# No SInE strategy applied
# Auto-Mode selected heuristic G_E___200_B02_F1_SE_CS_SP_PI_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/00055d0b1f0e491aac9357c25d6dcbb8.E_input', l32_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/00055d0b1f0e491aac9357c25d6dcbb8.E_input', t100_xboole_1)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/00055d0b1f0e491aac9357c25d6dcbb8.E_input', t17_xboole_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/00055d0b1f0e491aac9357c25d6dcbb8.E_input', idempotence_k3_xboole_0)).
fof(t114_zfmisc_1, conjecture, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X2,X1)=>(X1=k1_xboole_0|(X2=k1_xboole_0|X1=X2))), file('data/MPTP/proofs/00055d0b1f0e491aac9357c25d6dcbb8.E_input', t114_zfmisc_1)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/00055d0b1f0e491aac9357c25d6dcbb8.E_input', l54_zfmisc_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/00055d0b1f0e491aac9357c25d6dcbb8.E_input', t7_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/00055d0b1f0e491aac9357c25d6dcbb8.E_input', d3_tarski)).
fof(t112_xboole_1, axiom, ![X1, X2, X3]:k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k3_xboole_0(k5_xboole_0(X1,X3),X2), file('data/MPTP/proofs/00055d0b1f0e491aac9357c25d6dcbb8.E_input', t112_xboole_1)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/00055d0b1f0e491aac9357c25d6dcbb8.E_input', t83_xboole_1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/00055d0b1f0e491aac9357c25d6dcbb8.E_input', t4_xboole_0)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/00055d0b1f0e491aac9357c25d6dcbb8.E_input', d7_xboole_0)).
fof(t1_xboole_0, axiom, ![X1, X2, X3]:(r2_hidden(X1,k5_xboole_0(X2,X3))<=>~((r2_hidden(X1,X2)<=>r2_hidden(X1,X3)))), file('data/MPTP/proofs/00055d0b1f0e491aac9357c25d6dcbb8.E_input', t1_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/00055d0b1f0e491aac9357c25d6dcbb8.E_input', commutativity_k3_xboole_0)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/00055d0b1f0e491aac9357c25d6dcbb8.E_input', t48_xboole_1)).
fof(t5_boole, axiom, ![X1]:k5_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/00055d0b1f0e491aac9357c25d6dcbb8.E_input', t5_boole)).
fof(c_0_16, plain, ![X40, X41]:((k4_xboole_0(X40,X41)!=k1_xboole_0|r1_tarski(X40,X41))&(~r1_tarski(X40,X41)|k4_xboole_0(X40,X41)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
fof(c_0_17, plain, ![X9, X10]:k4_xboole_0(X9,X10)=k5_xboole_0(X9,k3_xboole_0(X9,X10)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_18, plain, ![X11, X12]:r1_tarski(k3_xboole_0(X11,X12),X11), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_19, plain, ![X26]:k3_xboole_0(X26,X26)=X26, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
fof(c_0_20, negated_conjecture, ~(![X1, X2]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X2,X1)=>(X1=k1_xboole_0|(X2=k1_xboole_0|X1=X2)))), inference(assume_negation,[status(cth)],[t114_zfmisc_1])).
fof(c_0_21, plain, ![X30, X31, X32, X33]:(((r2_hidden(X30,X32)|~r2_hidden(k4_tarski(X30,X31),k2_zfmisc_1(X32,X33)))&(r2_hidden(X31,X33)|~r2_hidden(k4_tarski(X30,X31),k2_zfmisc_1(X32,X33))))&(~r2_hidden(X30,X32)|~r2_hidden(X31,X33)|r2_hidden(k4_tarski(X30,X31),k2_zfmisc_1(X32,X33)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
fof(c_0_22, plain, ![X52]:(X52=k1_xboole_0|r2_hidden(esk6_1(X52),X52)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_23, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_27, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k2_zfmisc_1(esk2_0,esk1_0)&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&esk1_0!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).
cnf(c_0_28, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, plain, (X1=k1_xboole_0|r2_hidden(esk6_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_30, plain, ![X13, X14, X15, X16, X17]:((~r1_tarski(X13,X14)|(~r2_hidden(X15,X13)|r2_hidden(X15,X14)))&((r2_hidden(esk3_2(X16,X17),X16)|r1_tarski(X16,X17))&(~r2_hidden(esk3_2(X16,X17),X17)|r1_tarski(X16,X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_31, plain, ![X19, X20, X21]:k5_xboole_0(k3_xboole_0(X19,X20),k3_xboole_0(X21,X20))=k3_xboole_0(k5_xboole_0(X19,X21),X20), inference(variable_rename,[status(thm)],[t112_xboole_1])).
cnf(c_0_32, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=k1_xboole_0|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_33, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_34, plain, ![X44, X45]:((~r1_xboole_0(X44,X45)|k4_xboole_0(X44,X45)=X44)&(k4_xboole_0(X44,X45)!=X44|r1_xboole_0(X44,X45))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
fof(c_0_35, plain, ![X34, X35, X37, X38, X39]:((r1_xboole_0(X34,X35)|r2_hidden(esk5_2(X34,X35),k3_xboole_0(X34,X35)))&(~r2_hidden(X39,k3_xboole_0(X37,X38))|~r1_xboole_0(X37,X38))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
cnf(c_0_36, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_37, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k2_zfmisc_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_38, plain, (X1=k1_xboole_0|r2_hidden(k4_tarski(X2,esk6_1(X1)),k2_zfmisc_1(X3,X1))|~r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_39, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40, plain, (k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k3_xboole_0(k5_xboole_0(X1,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_26])).
fof(c_0_42, plain, ![X7, X8]:((~r1_xboole_0(X7,X8)|k3_xboole_0(X7,X8)=k1_xboole_0)&(k3_xboole_0(X7,X8)!=k1_xboole_0|r1_xboole_0(X7,X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
fof(c_0_43, plain, ![X49, X50, X51]:(((~r2_hidden(X49,X50)|~r2_hidden(X49,X51)|~r2_hidden(X49,k5_xboole_0(X50,X51)))&(r2_hidden(X49,X50)|r2_hidden(X49,X51)|~r2_hidden(X49,k5_xboole_0(X50,X51))))&((~r2_hidden(X49,X50)|r2_hidden(X49,X51)|r2_hidden(X49,k5_xboole_0(X50,X51)))&(~r2_hidden(X49,X51)|r2_hidden(X49,X50)|r2_hidden(X49,k5_xboole_0(X50,X51))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_0])])])).
cnf(c_0_44, plain, (r1_xboole_0(X1,X2)|k4_xboole_0(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_45, plain, ![X42, X43]:k3_xboole_0(X42,X43)=k3_xboole_0(X43,X42), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_46, plain, ![X27, X28]:k4_xboole_0(X27,k4_xboole_0(X27,X28))=k3_xboole_0(X27,X28), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_47, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_48, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(esk2_0,esk1_0))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_49, plain, (X1=k1_xboole_0|r2_hidden(k4_tarski(esk3_2(X2,X3),esk6_1(X1)),k2_zfmisc_1(X2,X1))|r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_50, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_51, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,X2))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_52, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_41])).
cnf(c_0_53, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_54, plain, (r2_hidden(X1,X3)|r2_hidden(X1,k5_xboole_0(X3,X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_55, plain, (r1_xboole_0(X1,X2)|k5_xboole_0(X1,k3_xboole_0(X1,X2))!=X1), inference(rw,[status(thm)],[c_0_44, c_0_24])).
cnf(c_0_56, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_57, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_58, plain, (~r2_hidden(X1,k5_xboole_0(k3_xboole_0(X2,X3),k3_xboole_0(X4,X3)))|~r1_xboole_0(k5_xboole_0(X2,X4),X3)), inference(spm,[status(thm)],[c_0_47, c_0_40])).
cnf(c_0_59, plain, (r1_tarski(X1,X2)|~r2_hidden(esk3_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_60, negated_conjecture, (r2_hidden(esk3_2(esk2_0,X1),esk1_0)|r1_tarski(esk2_0,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])).
cnf(c_0_61, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(k4_tarski(X2,X1),k2_zfmisc_1(esk2_0,esk1_0))), inference(spm,[status(thm)],[c_0_51, c_0_37])).
cnf(c_0_62, plain, (X1=k1_xboole_0|X2=k1_xboole_0|r2_hidden(k4_tarski(esk6_1(X1),esk6_1(X2)),k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_38, c_0_29])).
cnf(c_0_63, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_64, plain, (~r2_hidden(X1,k1_xboole_0)|~r1_xboole_0(k1_xboole_0,X2)), inference(spm,[status(thm)],[c_0_47, c_0_52])).
cnf(c_0_65, plain, (r1_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_53, c_0_52])).
cnf(c_0_66, plain, (r2_hidden(esk3_2(X1,X2),k5_xboole_0(X3,X1))|r2_hidden(esk3_2(X1,X2),X3)|r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_54, c_0_39])).
cnf(c_0_67, plain, (r1_xboole_0(X1,X2)|k5_xboole_0(X1,k3_xboole_0(X2,X1))!=X1), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_68, plain, (k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,k3_xboole_0(X1,X2)))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_25]), c_0_56])).
cnf(c_0_69, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2))))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_24]), c_0_24])).
cnf(c_0_70, plain, (~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X3,X2)))|~r1_xboole_0(k5_xboole_0(X2,X3),X2)), inference(spm,[status(thm)],[c_0_58, c_0_26])).
cnf(c_0_71, negated_conjecture, (r1_tarski(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_72, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_73, negated_conjecture, (r2_hidden(esk6_1(esk1_0),esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63]), c_0_50])).
cnf(c_0_74, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_65])])).
cnf(c_0_75, plain, (r2_hidden(esk3_2(X1,k5_xboole_0(X2,X1)),X2)|r1_tarski(X1,k5_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_59, c_0_66])).
cnf(c_0_76, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_77, plain, (r1_xboole_0(k3_xboole_0(X1,X2),X1)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_78, plain, (k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2)))=k5_xboole_0(X1,k3_xboole_0(X1,k3_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_69, c_0_69])).
cnf(c_0_79, plain, (k5_xboole_0(X1,k3_xboole_0(X2,X1))=k1_xboole_0|~r1_xboole_0(k5_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_70, c_0_29])).
cnf(c_0_80, negated_conjecture, (k5_xboole_0(esk2_0,k3_xboole_0(esk2_0,esk1_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_32, c_0_71])).
fof(c_0_81, plain, ![X29]:k5_xboole_0(X29,k1_xboole_0)=X29, inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_82, negated_conjecture, (r2_hidden(esk6_1(esk1_0),X1)|~r1_tarski(esk2_0,X1)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_83, plain, (r1_tarski(X1,k5_xboole_0(k1_xboole_0,X1))), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_84, plain, (k5_xboole_0(X1,X2)=k1_xboole_0|r2_hidden(esk6_1(k5_xboole_0(X1,X2)),X2)|r2_hidden(esk6_1(k5_xboole_0(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_76, c_0_29])).
cnf(c_0_85, plain, (r1_xboole_0(k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2)),k5_xboole_0(X1,X3))|k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_77, c_0_40])).
cnf(c_0_86, plain, (k5_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,k3_xboole_0(X1,X2))))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_69, c_0_78])).
cnf(c_0_87, negated_conjecture, (k5_xboole_0(esk2_0,k3_xboole_0(esk2_0,k3_xboole_0(esk2_0,esk1_0)))=k1_xboole_0), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_65])]), c_0_56])).
cnf(c_0_88, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_89, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_47, c_0_26])).
cnf(c_0_90, negated_conjecture, (r2_hidden(esk6_1(esk1_0),k5_xboole_0(k1_xboole_0,esk2_0))), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_91, plain, (r2_hidden(k4_tarski(X1,esk3_2(X2,X3)),k2_zfmisc_1(X4,X2))|r1_tarski(X2,X3)|~r2_hidden(X1,X4)), inference(spm,[status(thm)],[c_0_28, c_0_39])).
cnf(c_0_92, plain, (k5_xboole_0(k1_xboole_0,X1)=k1_xboole_0|r2_hidden(esk6_1(k5_xboole_0(k1_xboole_0,X1)),X1)), inference(spm,[status(thm)],[c_0_74, c_0_84])).
cnf(c_0_93, plain, (r1_xboole_0(k5_xboole_0(k1_xboole_0,k3_xboole_0(X1,X2)),k5_xboole_0(k1_xboole_0,X1))|k5_xboole_0(k1_xboole_0,k3_xboole_0(X1,X2))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_85, c_0_52])).
cnf(c_0_94, negated_conjecture, (k3_xboole_0(esk2_0,esk1_0)=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_88])).
cnf(c_0_95, negated_conjecture, (~r1_xboole_0(k5_xboole_0(k1_xboole_0,esk2_0),k5_xboole_0(k1_xboole_0,esk2_0))), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_96, plain, (k5_xboole_0(k1_xboole_0,X1)=k1_xboole_0|r2_hidden(k4_tarski(esk6_1(k5_xboole_0(k1_xboole_0,X1)),esk3_2(X2,X3)),k2_zfmisc_1(X1,X2))|r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_97, negated_conjecture, (k5_xboole_0(k1_xboole_0,esk2_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_95])).
cnf(c_0_98, negated_conjecture, (r2_hidden(esk3_2(esk1_0,X1),esk2_0)|r1_tarski(esk1_0,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_96]), c_0_97])).
cnf(c_0_99, plain, (k5_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,k3_xboole_0(X2,X1))))=k3_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_86, c_0_56])).
cnf(c_0_100, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_59, c_0_98])).
cnf(c_0_101, negated_conjecture, (k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,esk2_0))=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_94]), c_0_56]), c_0_94])).
cnf(c_0_102, negated_conjecture, (k5_xboole_0(esk1_0,esk2_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_100]), c_0_56]), c_0_94])).
cnf(c_0_103, negated_conjecture, (esk1_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_104, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_101, c_0_102]), c_0_88]), c_0_103]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 105
# Proof object clause steps            : 72
# Proof object formula steps           : 33
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 42
# Proof object simplifying inferences  : 26
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 32
# Processed clauses                    : 11937
# ...of these trivial                  : 601
# ...subsumed                          : 9049
# ...remaining for further processing  : 2287
# Other redundant clauses eliminated   : 69
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 23
# Backward-rewritten                   : 196
# Generated clauses                    : 173336
# ...of the previous two non-trivial   : 146897
# Contextual simplify-reflections      : 7
# Paramodulations                      : 173172
# Factorizations                       : 94
# Equation resolutions                 : 69
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2066
#    Positive orientable unit clauses  : 290
#    Positive unorientable unit clauses: 14
#    Negative unit clauses             : 64
#    Non-unit-clauses                  : 1698
# Current number of unprocessed clauses: 134142
# ...number of literals in the above   : 321431
# Current number of archived formulas  : 0
# Current number of archived clauses   : 221
# Clause-clause subsumption calls (NU) : 288785
# Rec. Clause-clause subsumption calls : 250568
# Non-unit clause-clause subsumptions  : 4878
# Unit Clause-clause subsumption calls : 19507
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1256
# BW rewrite match successes           : 221
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4183828

# -------------------------------------------------
# User time                : 1.813 s
# System time              : 0.037 s
# Total time               : 1.850 s
# Maximum resident set size: 119108 pages
