# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t53_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>(![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>r1_yellow_0(X1,X2))=>v3_lattice3(X1))), file('data/MPTP/proofs/005c69f0364744718d955ae00ad00c6c.E_input', t53_yellow_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/005c69f0364744718d955ae00ad00c6c.E_input', t3_subset)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/005c69f0364744718d955ae00ad00c6c.E_input', t17_xboole_1)).
fof(d12_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v3_lattice3(X1)<=>![X2]:?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4))))))), file('data/MPTP/proofs/005c69f0364744718d955ae00ad00c6c.E_input', d12_lattice3)).
fof(d7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X5)=>r1_orders_2(X1,X4,X5))))=>X4=X3))))))), file('data/MPTP/proofs/005c69f0364744718d955ae00ad00c6c.E_input', d7_yellow_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/005c69f0364744718d955ae00ad00c6c.E_input', commutativity_k3_xboole_0)).
fof(t50_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((r1_yellow_0(X1,X2)=>r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))&((r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))=>r1_yellow_0(X1,X2))&((r2_yellow_0(X1,X2)=>r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))&(r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))=>r2_yellow_0(X1,X2)))))), file('data/MPTP/proofs/005c69f0364744718d955ae00ad00c6c.E_input', t50_yellow_0)).
fof(c_0_7, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>(![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>r1_yellow_0(X1,X2))=>v3_lattice3(X1)))), inference(assume_negation,[status(cth)],[t53_yellow_0])).
fof(c_0_8, plain, ![X8, X9]:((~m1_subset_1(X8,k1_zfmisc_1(X9))|r1_tarski(X8,X9))&(~r1_tarski(X8,X9)|m1_subset_1(X8,k1_zfmisc_1(X9)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_9, plain, ![X10, X11]:r1_tarski(k3_xboole_0(X10,X11),X10), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_10, plain, ![X12, X13, X15, X17]:(((m1_subset_1(esk2_2(X12,X13),u1_struct_0(X12))|~v3_lattice3(X12)|~l1_orders_2(X12))&((r2_lattice3(X12,X13,esk2_2(X12,X13))|~v3_lattice3(X12)|~l1_orders_2(X12))&(~m1_subset_1(X15,u1_struct_0(X12))|(~r2_lattice3(X12,X13,X15)|r1_orders_2(X12,esk2_2(X12,X13),X15))|~v3_lattice3(X12)|~l1_orders_2(X12))))&((m1_subset_1(esk4_2(X12,X17),u1_struct_0(X12))|~r2_lattice3(X12,esk3_1(X12),X17)|~m1_subset_1(X17,u1_struct_0(X12))|v3_lattice3(X12)|~l1_orders_2(X12))&((r2_lattice3(X12,esk3_1(X12),esk4_2(X12,X17))|~r2_lattice3(X12,esk3_1(X12),X17)|~m1_subset_1(X17,u1_struct_0(X12))|v3_lattice3(X12)|~l1_orders_2(X12))&(~r1_orders_2(X12,X17,esk4_2(X12,X17))|~r2_lattice3(X12,esk3_1(X12),X17)|~m1_subset_1(X17,u1_struct_0(X12))|v3_lattice3(X12)|~l1_orders_2(X12))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_lattice3])])])])])).
fof(c_0_11, plain, ![X19, X20, X22, X23, X25, X26, X29]:(((m1_subset_1(esk5_2(X19,X20),u1_struct_0(X19))|~r1_yellow_0(X19,X20)|~l1_orders_2(X19))&((r2_lattice3(X19,X20,esk5_2(X19,X20))|~r1_yellow_0(X19,X20)|~l1_orders_2(X19))&((~m1_subset_1(X22,u1_struct_0(X19))|(~r2_lattice3(X19,X20,X22)|r1_orders_2(X19,esk5_2(X19,X20),X22))|~r1_yellow_0(X19,X20)|~l1_orders_2(X19))&((m1_subset_1(esk6_3(X19,X20,X23),u1_struct_0(X19))|~r2_lattice3(X19,X20,X23)|X23=esk5_2(X19,X20)|~m1_subset_1(X23,u1_struct_0(X19))|~r1_yellow_0(X19,X20)|~l1_orders_2(X19))&((r2_lattice3(X19,X20,esk6_3(X19,X20,X23))|~r2_lattice3(X19,X20,X23)|X23=esk5_2(X19,X20)|~m1_subset_1(X23,u1_struct_0(X19))|~r1_yellow_0(X19,X20)|~l1_orders_2(X19))&(~r1_orders_2(X19,X23,esk6_3(X19,X20,X23))|~r2_lattice3(X19,X20,X23)|X23=esk5_2(X19,X20)|~m1_subset_1(X23,u1_struct_0(X19))|~r1_yellow_0(X19,X20)|~l1_orders_2(X19)))))))&(((m1_subset_1(esk8_3(X19,X25,X26),u1_struct_0(X19))|m1_subset_1(esk7_3(X19,X25,X26),u1_struct_0(X19))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19))&(((r2_lattice3(X19,X25,esk8_3(X19,X25,X26))|m1_subset_1(esk7_3(X19,X25,X26),u1_struct_0(X19))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19))&(~m1_subset_1(X29,u1_struct_0(X19))|(~r2_lattice3(X19,X25,X29)|r1_orders_2(X19,esk8_3(X19,X25,X26),X29))|m1_subset_1(esk7_3(X19,X25,X26),u1_struct_0(X19))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19)))&(esk8_3(X19,X25,X26)!=X26|m1_subset_1(esk7_3(X19,X25,X26),u1_struct_0(X19))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19))))&(((m1_subset_1(esk8_3(X19,X25,X26),u1_struct_0(X19))|r2_lattice3(X19,X25,esk7_3(X19,X25,X26))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19))&(((r2_lattice3(X19,X25,esk8_3(X19,X25,X26))|r2_lattice3(X19,X25,esk7_3(X19,X25,X26))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19))&(~m1_subset_1(X29,u1_struct_0(X19))|(~r2_lattice3(X19,X25,X29)|r1_orders_2(X19,esk8_3(X19,X25,X26),X29))|r2_lattice3(X19,X25,esk7_3(X19,X25,X26))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19)))&(esk8_3(X19,X25,X26)!=X26|r2_lattice3(X19,X25,esk7_3(X19,X25,X26))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19))))&((m1_subset_1(esk8_3(X19,X25,X26),u1_struct_0(X19))|~r1_orders_2(X19,X26,esk7_3(X19,X25,X26))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19))&(((r2_lattice3(X19,X25,esk8_3(X19,X25,X26))|~r1_orders_2(X19,X26,esk7_3(X19,X25,X26))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19))&(~m1_subset_1(X29,u1_struct_0(X19))|(~r2_lattice3(X19,X25,X29)|r1_orders_2(X19,esk8_3(X19,X25,X26),X29))|~r1_orders_2(X19,X26,esk7_3(X19,X25,X26))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19)))&(esk8_3(X19,X25,X26)!=X26|~r1_orders_2(X19,X26,esk7_3(X19,X25,X26))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_yellow_0])])])])])])).
fof(c_0_12, negated_conjecture, ![X7]:((~v2_struct_0(esk1_0)&l1_orders_2(esk1_0))&((~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_yellow_0(esk1_0,X7))&~v3_lattice3(esk1_0))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_7])])])])])).
cnf(c_0_13, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (v3_lattice3(X1)|~r1_orders_2(X1,X2,esk4_2(X1,X2))|~r2_lattice3(X1,esk3_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (r1_orders_2(X2,esk5_2(X2,X3),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r1_yellow_0(X2,X3)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, negated_conjecture, (r1_yellow_0(esk1_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, plain, (m1_subset_1(k3_xboole_0(X1,X2),k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_13, c_0_14])).
fof(c_0_20, plain, ![X30, X31]:k3_xboole_0(X30,X31)=k3_xboole_0(X31,X30), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_21, plain, (v3_lattice3(X1)|~r2_lattice3(X1,X2,esk4_2(X1,esk5_2(X1,X2)))|~r2_lattice3(X1,esk3_1(X1),esk5_2(X1,X2))|~r1_yellow_0(X1,X2)|~m1_subset_1(esk4_2(X1,esk5_2(X1,X2)),u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])).
cnf(c_0_22, plain, (r2_lattice3(X1,esk3_1(X1),esk4_2(X1,X2))|v3_lattice3(X1)|~r2_lattice3(X1,esk3_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_23, plain, (r2_lattice3(X1,X2,esk5_2(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_24, plain, ![X32, X33]:((~r1_yellow_0(X32,X33)|r1_yellow_0(X32,k3_xboole_0(X33,u1_struct_0(X32)))|(v2_struct_0(X32)|~l1_orders_2(X32)))&((~r1_yellow_0(X32,k3_xboole_0(X33,u1_struct_0(X32)))|r1_yellow_0(X32,X33)|(v2_struct_0(X32)|~l1_orders_2(X32)))&((~r2_yellow_0(X32,X33)|r2_yellow_0(X32,k3_xboole_0(X33,u1_struct_0(X32)))|(v2_struct_0(X32)|~l1_orders_2(X32)))&(~r2_yellow_0(X32,k3_xboole_0(X33,u1_struct_0(X32)))|r2_yellow_0(X32,X33)|(v2_struct_0(X32)|~l1_orders_2(X32)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_yellow_0])])])])])).
cnf(c_0_25, negated_conjecture, (r1_yellow_0(esk1_0,k3_xboole_0(u1_struct_0(esk1_0),X1))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_26, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (v3_lattice3(X1)|~r1_yellow_0(X1,esk3_1(X1))|~m1_subset_1(esk4_2(X1,esk5_2(X1,esk3_1(X1))),u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_17]), c_0_23])).
cnf(c_0_28, plain, (m1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|v3_lattice3(X1)|~r2_lattice3(X1,esk3_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_29, plain, (r1_yellow_0(X1,X2)|v2_struct_0(X1)|~r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, negated_conjecture, (r1_yellow_0(esk1_0,k3_xboole_0(X1,u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_31, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_32, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_33, plain, (v3_lattice3(X1)|~r1_yellow_0(X1,esk3_1(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_17]), c_0_23])).
cnf(c_0_34, negated_conjecture, (r1_yellow_0(esk1_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])]), c_0_32])).
cnf(c_0_35, negated_conjecture, (~v3_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_36, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_31])]), c_0_35]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 37
# Proof object clause steps            : 22
# Proof object formula steps           : 15
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 36
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 36
# Processed clauses                    : 95
# ...of these trivial                  : 4
# ...subsumed                          : 9
# ...remaining for further processing  : 82
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 3
# Generated clauses                    : 55
# ...of the previous two non-trivial   : 39
# Contextual simplify-reflections      : 5
# Paramodulations                      : 55
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 42
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 33
# Current number of unprocessed clauses: 15
# ...number of literals in the above   : 134
# Current number of archived formulas  : 0
# Current number of archived clauses   : 40
# Clause-clause subsumption calls (NU) : 1014
# Rec. Clause-clause subsumption calls : 62
# Non-unit clause-clause subsumptions  : 15
# Unit Clause-clause subsumption calls : 3
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4467

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28692 pages
