# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t1_tex_2, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((~(v1_xboole_0(X2))&v1_zfmisc_1(X2))=>(r1_tarski(X1,X2)=>X1=X2))), file('data/MPTP/proofs/00840c1073db4d1a9a44d8c428f903fe.E_input', t1_tex_2)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/00840c1073db4d1a9a44d8c428f903fe.E_input', d5_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/00840c1073db4d1a9a44d8c428f903fe.E_input', d3_tarski)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/00840c1073db4d1a9a44d8c428f903fe.E_input', t3_xboole_0)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/00840c1073db4d1a9a44d8c428f903fe.E_input', t1_xboole_1)).
fof(t10_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(X1,k2_xboole_0(X3,X2))), file('data/MPTP/proofs/00840c1073db4d1a9a44d8c428f903fe.E_input', t10_xboole_1)).
fof(d3_xboole_0, axiom, ![X1, X2, X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2)))), file('data/MPTP/proofs/00840c1073db4d1a9a44d8c428f903fe.E_input', d3_xboole_0)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/00840c1073db4d1a9a44d8c428f903fe.E_input', t1_boole)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/00840c1073db4d1a9a44d8c428f903fe.E_input', commutativity_k2_xboole_0)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/00840c1073db4d1a9a44d8c428f903fe.E_input', t83_xboole_1)).
fof(t69_xboole_1, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>~((r1_tarski(X2,X1)&r1_xboole_0(X2,X1)))), file('data/MPTP/proofs/00840c1073db4d1a9a44d8c428f903fe.E_input', t69_xboole_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/00840c1073db4d1a9a44d8c428f903fe.E_input', d1_xboole_0)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/00840c1073db4d1a9a44d8c428f903fe.E_input', t12_xboole_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/00840c1073db4d1a9a44d8c428f903fe.E_input', l32_xboole_1)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/00840c1073db4d1a9a44d8c428f903fe.E_input', redefinition_k6_domain_1)).
fof(d1_tex_2, axiom, ![X1]:(~(v1_xboole_0(X1))=>(v1_zfmisc_1(X1)<=>?[X2]:(m1_subset_1(X2,X1)&X1=k6_domain_1(X1,X2)))), file('data/MPTP/proofs/00840c1073db4d1a9a44d8c428f903fe.E_input', d1_tex_2)).
fof(l3_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(X1,k1_tarski(X2))<=>(X1=k1_xboole_0|X1=k1_tarski(X2))), file('data/MPTP/proofs/00840c1073db4d1a9a44d8c428f903fe.E_input', l3_zfmisc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/00840c1073db4d1a9a44d8c428f903fe.E_input', d10_xboole_0)).
fof(c_0_18, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:((~(v1_xboole_0(X2))&v1_zfmisc_1(X2))=>(r1_tarski(X1,X2)=>X1=X2)))), inference(assume_negation,[status(cth)],[t1_tex_2])).
fof(c_0_19, plain, ![X7, X8, X9, X10, X11, X12, X13, X14]:((((r2_hidden(X10,X7)|~r2_hidden(X10,X9)|X9!=k4_xboole_0(X7,X8))&(~r2_hidden(X10,X8)|~r2_hidden(X10,X9)|X9!=k4_xboole_0(X7,X8)))&(~r2_hidden(X11,X7)|r2_hidden(X11,X8)|r2_hidden(X11,X9)|X9!=k4_xboole_0(X7,X8)))&((~r2_hidden(esk3_3(X12,X13,X14),X14)|(~r2_hidden(esk3_3(X12,X13,X14),X12)|r2_hidden(esk3_3(X12,X13,X14),X13))|X14=k4_xboole_0(X12,X13))&((r2_hidden(esk3_3(X12,X13,X14),X12)|r2_hidden(esk3_3(X12,X13,X14),X14)|X14=k4_xboole_0(X12,X13))&(~r2_hidden(esk3_3(X12,X13,X14),X13)|r2_hidden(esk3_3(X12,X13,X14),X14)|X14=k4_xboole_0(X12,X13))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_20, plain, ![X16, X17, X18, X19, X20]:((~r1_tarski(X16,X17)|(~r2_hidden(X18,X16)|r2_hidden(X18,X17)))&((r2_hidden(esk4_2(X19,X20),X19)|r1_tarski(X19,X20))&(~r2_hidden(esk4_2(X19,X20),X20)|r1_tarski(X19,X20)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_21, negated_conjecture, (~v1_xboole_0(esk1_0)&((~v1_xboole_0(esk2_0)&v1_zfmisc_1(esk2_0))&(r1_tarski(esk1_0,esk2_0)&esk1_0!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_18])])])])).
cnf(c_0_22, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_23, plain, ![X22, X23, X25, X26, X27]:(((r2_hidden(esk5_2(X22,X23),X22)|r1_xboole_0(X22,X23))&(r2_hidden(esk5_2(X22,X23),X23)|r1_xboole_0(X22,X23)))&(~r2_hidden(X27,X25)|~r2_hidden(X27,X26)|~r1_xboole_0(X25,X26))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_24, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_26, plain, ![X28, X29, X30]:(~r1_tarski(X28,X29)|~r1_tarski(X29,X30)|r1_tarski(X28,X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_27, plain, ![X31, X32, X33]:(~r1_tarski(X31,X32)|r1_tarski(X31,k2_xboole_0(X33,X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_xboole_1])])).
cnf(c_0_28, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_29, plain, (r2_hidden(esk5_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_32, plain, ![X34, X35, X36, X37, X38, X39, X40, X41]:(((~r2_hidden(X37,X36)|(r2_hidden(X37,X34)|r2_hidden(X37,X35))|X36!=k2_xboole_0(X34,X35))&((~r2_hidden(X38,X34)|r2_hidden(X38,X36)|X36!=k2_xboole_0(X34,X35))&(~r2_hidden(X38,X35)|r2_hidden(X38,X36)|X36!=k2_xboole_0(X34,X35))))&(((~r2_hidden(esk6_3(X39,X40,X41),X39)|~r2_hidden(esk6_3(X39,X40,X41),X41)|X41=k2_xboole_0(X39,X40))&(~r2_hidden(esk6_3(X39,X40,X41),X40)|~r2_hidden(esk6_3(X39,X40,X41),X41)|X41=k2_xboole_0(X39,X40)))&(r2_hidden(esk6_3(X39,X40,X41),X41)|(r2_hidden(esk6_3(X39,X40,X41),X39)|r2_hidden(esk6_3(X39,X40,X41),X40))|X41=k2_xboole_0(X39,X40)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])).
cnf(c_0_33, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (r1_tarski(X1,k2_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_35, plain, ![X43]:k2_xboole_0(X43,k1_xboole_0)=X43, inference(variable_rename,[status(thm)],[t1_boole])).
fof(c_0_36, plain, ![X44, X45]:k2_xboole_0(X44,X45)=k2_xboole_0(X45,X44), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_37, plain, ![X46, X47]:((~r1_xboole_0(X46,X47)|k4_xboole_0(X46,X47)=X46)&(k4_xboole_0(X46,X47)!=X46|r1_xboole_0(X46,X47))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
cnf(c_0_38, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r2_hidden(esk5_2(X1,k4_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_39, negated_conjecture, (r1_xboole_0(esk1_0,X1)|r2_hidden(esk5_2(esk1_0,X1),esk2_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
fof(c_0_40, plain, ![X48, X49]:(v1_xboole_0(X49)|(~r1_tarski(X49,X48)|~r1_xboole_0(X49,X48))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t69_xboole_1])])])).
fof(c_0_41, plain, ![X50, X51, X52]:((~v1_xboole_0(X50)|~r2_hidden(X51,X50))&(r2_hidden(esk7_1(X52),X52)|v1_xboole_0(X52))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_42, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_43, plain, ![X54, X55]:(~r1_tarski(X54,X55)|k2_xboole_0(X54,X55)=X55), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_44, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,X4)|~r1_tarski(X4,X3)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_45, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_46, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_47, plain, ![X56, X57]:((k4_xboole_0(X56,X57)!=k1_xboole_0|r1_tarski(X56,X57))&(~r1_tarski(X56,X57)|k4_xboole_0(X56,X57)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_48, plain, (k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_49, negated_conjecture, (r1_xboole_0(esk1_0,k4_xboole_0(X1,esk2_0))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_50, plain, (v1_xboole_0(X1)|~r1_tarski(X1,X2)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_51, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_52, plain, ![X58, X59]:(v1_xboole_0(X58)|~m1_subset_1(X59,X58)|k6_domain_1(X58,X59)=k1_tarski(X59)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_53, plain, ![X60, X62]:(((m1_subset_1(esk8_1(X60),X60)|~v1_zfmisc_1(X60)|v1_xboole_0(X60))&(X60=k6_domain_1(X60,esk8_1(X60))|~v1_zfmisc_1(X60)|v1_xboole_0(X60)))&(~m1_subset_1(X62,X60)|X60!=k6_domain_1(X60,X62)|v1_zfmisc_1(X60)|v1_xboole_0(X60))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_tex_2])])])])])])).
cnf(c_0_54, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_55, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_56, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_57, plain, ![X63, X64]:((~r1_tarski(X63,k1_tarski(X64))|(X63=k1_xboole_0|X63=k1_tarski(X64)))&((X63!=k1_xboole_0|r1_tarski(X63,k1_tarski(X64)))&(X63!=k1_tarski(X64)|r1_tarski(X63,k1_tarski(X64))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_zfmisc_1])])])).
cnf(c_0_58, negated_conjecture, (r1_tarski(esk1_0,k2_xboole_0(X1,X2))|~r1_tarski(esk2_0,X2)), inference(spm,[status(thm)],[c_0_44, c_0_25])).
cnf(c_0_59, plain, (k2_xboole_0(k1_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_60, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_61, negated_conjecture, (k4_xboole_0(esk1_0,k4_xboole_0(X1,esk2_0))=esk1_0), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_62, negated_conjecture, (~r1_tarski(esk1_0,k4_xboole_0(X1,esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_49]), c_0_51])).
cnf(c_0_63, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_64, plain, (m1_subset_1(esk8_1(X1),X1)|v1_xboole_0(X1)|~v1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_65, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(k2_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_66, plain, (k2_xboole_0(X1,X2)=X1|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_46, c_0_56])).
cnf(c_0_67, plain, (r2_hidden(esk7_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_68, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|~r1_tarski(X1,k1_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_69, negated_conjecture, (r1_tarski(esk1_0,X1)|~r1_tarski(esk2_0,X1)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_70, negated_conjecture, (esk1_0!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62])).
cnf(c_0_71, plain, (X1=k6_domain_1(X1,esk8_1(X1))|v1_xboole_0(X1)|~v1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_72, plain, (k6_domain_1(X1,esk8_1(X1))=k1_tarski(esk8_1(X1))|v1_xboole_0(X1)|~v1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_73, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X3)|~v1_xboole_0(X3)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_74, negated_conjecture, (r2_hidden(esk7_1(esk1_0),esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_67]), c_0_51])).
fof(c_0_75, plain, ![X65, X66]:(((r1_tarski(X65,X66)|X65!=X66)&(r1_tarski(X66,X65)|X65!=X66))&(~r1_tarski(X65,X66)|~r1_tarski(X66,X65)|X65=X66)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_76, negated_conjecture, (k1_tarski(X1)=esk1_0|~r1_tarski(esk2_0,k1_tarski(X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70])).
cnf(c_0_77, plain, (k1_tarski(esk8_1(X1))=X1|v1_xboole_0(X1)|~v1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_78, negated_conjecture, (~r1_tarski(esk2_0,X1)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_79, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_80, negated_conjecture, (X1=esk1_0|~r1_tarski(esk2_0,X1)|~v1_zfmisc_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78])).
cnf(c_0_81, negated_conjecture, (v1_zfmisc_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_82, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_79])).
cnf(c_0_83, negated_conjecture, (esk1_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_84, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_82])]), c_0_83]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 85
# Proof object clause steps            : 48
# Proof object formula steps           : 37
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 45
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 45
# Processed clauses                    : 1541
# ...of these trivial                  : 17
# ...subsumed                          : 1115
# ...remaining for further processing  : 409
# Other redundant clauses eliminated   : 28
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 13
# Backward-rewritten                   : 8
# Generated clauses                    : 4215
# ...of the previous two non-trivial   : 3726
# Contextual simplify-reflections      : 1
# Paramodulations                      : 4099
# Factorizations                       : 88
# Equation resolutions                 : 28
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 335
#    Positive orientable unit clauses  : 39
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 36
#    Non-unit-clauses                  : 259
# Current number of unprocessed clauses: 2253
# ...number of literals in the above   : 7776
# Current number of archived formulas  : 0
# Current number of archived clauses   : 64
# Clause-clause subsumption calls (NU) : 25321
# Rec. Clause-clause subsumption calls : 14266
# Non-unit clause-clause subsumptions  : 683
# Unit Clause-clause subsumption calls : 2998
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 47
# BW rewrite match successes           : 14
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 46329

# -------------------------------------------------
# User time                : 0.041 s
# System time              : 0.007 s
# Total time               : 0.048 s
# Maximum resident set size: 28632 pages
