# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t28_zfmisc_1, conjecture, ![X1, X2, X3, X4]:~((r1_tarski(k2_tarski(X1,X2),k2_tarski(X3,X4))&(~(X1=X3)&~(X1=X4)))), file('data/MPTP/proofs/00863e3bcfc4493185ca31a0526b1385.E_input', t28_zfmisc_1)).
fof(l45_zfmisc_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_tarski(X2,X3))<=>~((~(X1=k1_xboole_0)&(~(X1=k1_tarski(X2))&(~(X1=k1_tarski(X3))&~(X1=k2_tarski(X2,X3))))))), file('data/MPTP/proofs/00863e3bcfc4493185ca31a0526b1385.E_input', l45_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/00863e3bcfc4493185ca31a0526b1385.E_input', t70_enumset1)).
fof(t89_enumset1, axiom, ![X1, X2, X3]:k5_enumset1(X1,X1,X1,X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/00863e3bcfc4493185ca31a0526b1385.E_input', t89_enumset1)).
fof(t70_xboole_1, axiom, ![X1, X2, X3]:(~((~(r1_xboole_0(X1,k2_xboole_0(X2,X3)))&(r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3))))&~((~((r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3)))&r1_xboole_0(X1,k2_xboole_0(X2,X3))))), file('data/MPTP/proofs/00863e3bcfc4493185ca31a0526b1385.E_input', t70_xboole_1)).
fof(t51_xboole_1, axiom, ![X1, X2]:k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/00863e3bcfc4493185ca31a0526b1385.E_input', t51_xboole_1)).
fof(t25_zfmisc_1, axiom, ![X1, X2, X3]:~((r1_tarski(k1_tarski(X1),k2_tarski(X2,X3))&(~(X1=X2)&~(X1=X3)))), file('data/MPTP/proofs/00863e3bcfc4493185ca31a0526b1385.E_input', t25_zfmisc_1)).
fof(t81_xboole_1, axiom, ![X1, X2, X3]:(r1_xboole_0(X1,k4_xboole_0(X2,X3))=>r1_xboole_0(X2,k4_xboole_0(X1,X3))), file('data/MPTP/proofs/00863e3bcfc4493185ca31a0526b1385.E_input', t81_xboole_1)).
fof(t66_xboole_1, axiom, ![X1]:(~((~(r1_xboole_0(X1,X1))&X1=k1_xboole_0))&~((~(X1=k1_xboole_0)&r1_xboole_0(X1,X1)))), file('data/MPTP/proofs/00863e3bcfc4493185ca31a0526b1385.E_input', t66_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/00863e3bcfc4493185ca31a0526b1385.E_input', d10_xboole_0)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/00863e3bcfc4493185ca31a0526b1385.E_input', l32_xboole_1)).
fof(t9_zfmisc_1, axiom, ![X1, X2, X3]:(k1_tarski(X1)=k2_tarski(X2,X3)=>X2=X3), file('data/MPTP/proofs/00863e3bcfc4493185ca31a0526b1385.E_input', t9_zfmisc_1)).
fof(t18_zfmisc_1, axiom, ![X1, X2]:(k3_xboole_0(k1_tarski(X1),k1_tarski(X2))=k1_tarski(X1)=>X1=X2), file('data/MPTP/proofs/00863e3bcfc4493185ca31a0526b1385.E_input', t18_zfmisc_1)).
fof(t12_zfmisc_1, axiom, ![X1, X2]:r1_tarski(k1_tarski(X1),k2_tarski(X1,X2)), file('data/MPTP/proofs/00863e3bcfc4493185ca31a0526b1385.E_input', t12_zfmisc_1)).
fof(c_0_14, negated_conjecture, ~(![X1, X2, X3, X4]:~((r1_tarski(k2_tarski(X1,X2),k2_tarski(X3,X4))&(~(X1=X3)&~(X1=X4))))), inference(assume_negation,[status(cth)],[t28_zfmisc_1])).
fof(c_0_15, plain, ![X11, X12, X13]:((~r1_tarski(X11,k2_tarski(X12,X13))|(X11=k1_xboole_0|(X11=k1_tarski(X12)|(X11=k1_tarski(X13)|X11=k2_tarski(X12,X13)))))&((X11!=k1_xboole_0|r1_tarski(X11,k2_tarski(X12,X13)))&((X11!=k1_tarski(X12)|r1_tarski(X11,k2_tarski(X12,X13)))&((X11!=k1_tarski(X13)|r1_tarski(X11,k2_tarski(X12,X13)))&(X11!=k2_tarski(X12,X13)|r1_tarski(X11,k2_tarski(X12,X13))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l45_zfmisc_1])])])])).
fof(c_0_16, plain, ![X17, X18]:k1_enumset1(X17,X17,X18)=k2_tarski(X17,X18), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_17, plain, ![X24, X25, X26]:k5_enumset1(X24,X24,X24,X24,X24,X25,X26)=k1_enumset1(X24,X25,X26), inference(variable_rename,[status(thm)],[t89_enumset1])).
fof(c_0_18, negated_conjecture, (r1_tarski(k2_tarski(esk1_0,esk2_0),k2_tarski(esk3_0,esk4_0))&(esk1_0!=esk3_0&esk1_0!=esk4_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])).
fof(c_0_19, plain, ![X27, X28, X29]:((r1_xboole_0(X27,k2_xboole_0(X28,X29))|(~r1_xboole_0(X27,X28)|~r1_xboole_0(X27,X29)))&((r1_xboole_0(X27,X28)|~r1_xboole_0(X27,k2_xboole_0(X28,X29)))&(r1_xboole_0(X27,X29)|~r1_xboole_0(X27,k2_xboole_0(X28,X29))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t70_xboole_1])])])])).
fof(c_0_20, plain, ![X38, X39]:k2_xboole_0(k3_xboole_0(X38,X39),k4_xboole_0(X38,X39))=X38, inference(variable_rename,[status(thm)],[t51_xboole_1])).
fof(c_0_21, plain, ![X21, X22, X23]:(~r1_tarski(k1_tarski(X21),k2_tarski(X22,X23))|(X21=X22|X21=X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t25_zfmisc_1])])])).
cnf(c_0_22, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|X1=k1_tarski(X3)|X1=k2_tarski(X2,X3)|~r1_tarski(X1,k2_tarski(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (k5_enumset1(X1,X1,X1,X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, negated_conjecture, (r1_tarski(k2_tarski(esk1_0,esk2_0),k2_tarski(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_26, plain, ![X14, X15, X16]:(~r1_xboole_0(X14,k4_xboole_0(X15,X16))|r1_xboole_0(X15,k4_xboole_0(X14,X16))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t81_xboole_1])])).
cnf(c_0_27, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X1,k2_xboole_0(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, plain, (k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_29, plain, ![X34]:((r1_xboole_0(X34,X34)|X34!=k1_xboole_0)&(X34=k1_xboole_0|~r1_xboole_0(X34,X34))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t66_xboole_1])])])).
cnf(c_0_30, plain, (X1=X2|X1=X3|~r1_tarski(k1_tarski(X1),k2_tarski(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, plain, (X1=k1_xboole_0|X1=k1_tarski(X3)|X1=k1_tarski(X2)|X1=k5_enumset1(X2,X2,X2,X2,X2,X2,X3)|~r1_tarski(X1,k5_enumset1(X2,X2,X2,X2,X2,X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_23]), c_0_23]), c_0_24]), c_0_24])).
cnf(c_0_32, negated_conjecture, (r1_tarski(k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0),k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_23]), c_0_23]), c_0_24]), c_0_24])).
fof(c_0_33, plain, ![X9, X10]:(((r1_tarski(X9,X10)|X9!=X10)&(r1_tarski(X10,X9)|X9!=X10))&(~r1_tarski(X9,X10)|~r1_tarski(X10,X9)|X9=X10)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_34, plain, (r1_xboole_0(X2,k4_xboole_0(X1,X3))|~r1_xboole_0(X1,k4_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_36, plain, (r1_xboole_0(X1,X1)|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, plain, (r1_tarski(X1,k2_tarski(X3,X2))|X1!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_38, plain, (X1=X3|X1=X2|~r1_tarski(k1_tarski(X1),k5_enumset1(X2,X2,X2,X2,X2,X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_23]), c_0_24])).
cnf(c_0_39, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_40, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_41, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_42, plain, (r1_xboole_0(k1_xboole_0,k1_xboole_0)), inference(er,[status(thm)],[c_0_36])).
cnf(c_0_43, plain, (r1_tarski(X1,k5_enumset1(X3,X3,X3,X3,X3,X3,X2))|X1!=k1_tarski(X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_23]), c_0_24])).
cnf(c_0_44, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|X1=esk1_0|X1=esk2_0|~r1_tarski(k1_tarski(X1),k1_tarski(esk4_0))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_45, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_46, negated_conjecture, (esk1_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_47, plain, (r1_xboole_0(k1_xboole_0,k4_xboole_0(k1_xboole_0,X1))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_48, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_49, plain, (r1_tarski(k1_tarski(X1),k5_enumset1(X2,X2,X2,X2,X2,X2,X1))), inference(er,[status(thm)],[c_0_43])).
cnf(c_0_50, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|esk2_0=esk4_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])).
fof(c_0_51, plain, ![X32, X33]:((k4_xboole_0(X32,X33)!=k1_xboole_0|r1_tarski(X32,X33))&(~r1_tarski(X32,X33)|k4_xboole_0(X32,X33)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_52, plain, (X1=k1_xboole_0|~r1_xboole_0(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_53, plain, (r1_xboole_0(k4_xboole_0(k1_xboole_0,X1),k4_xboole_0(k1_xboole_0,X2))), inference(spm,[status(thm)],[c_0_41, c_0_47])).
cnf(c_0_54, plain, (r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_48, c_0_28])).
cnf(c_0_55, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_56, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|esk2_0=esk4_0|r1_tarski(k1_tarski(esk2_0),k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_57, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_58, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_59, plain, (r1_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X3,X4))|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_41, c_0_54])).
fof(c_0_60, plain, ![X35, X36, X37]:(k1_tarski(X35)!=k2_tarski(X36,X37)|X36=X37), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_zfmisc_1])])).
fof(c_0_61, plain, ![X30, X31]:(k3_xboole_0(k1_tarski(X30),k1_tarski(X31))!=k1_tarski(X30)|X30=X31), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t18_zfmisc_1])])).
cnf(c_0_62, plain, (k5_enumset1(X1,X1,X1,X1,X1,X1,X2)=k1_tarski(X2)|~r1_tarski(k5_enumset1(X1,X1,X1,X1,X1,X1,X2),k1_tarski(X2))), inference(spm,[status(thm)],[c_0_55, c_0_49])).
cnf(c_0_63, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|esk2_0=esk3_0|esk2_0=esk4_0), inference(spm,[status(thm)],[c_0_38, c_0_56])).
cnf(c_0_64, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_65, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(k3_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_52, c_0_54])).
cnf(c_0_66, plain, (r1_xboole_0(k3_xboole_0(k1_xboole_0,X1),k1_xboole_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_42]), c_0_58])).
cnf(c_0_67, plain, (X2=X3|k1_tarski(X1)!=k2_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_68, plain, (X1=X2|k3_xboole_0(k1_tarski(X1),k1_tarski(X2))!=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_69, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|k1_tarski(esk2_0)=k1_xboole_0|esk2_0=esk4_0|esk2_0=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64])])).
cnf(c_0_70, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_65, c_0_66])).
fof(c_0_71, plain, ![X19, X20]:r1_tarski(k1_tarski(X19),k2_tarski(X19,X20)), inference(variable_rename,[status(thm)],[t12_zfmisc_1])).
cnf(c_0_72, plain, (X2=X3|k1_tarski(X1)!=k5_enumset1(X2,X2,X2,X2,X2,X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_23]), c_0_24])).
cnf(c_0_73, plain, (r1_tarski(X1,k2_tarski(X2,X3))|X1!=k2_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_74, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|esk2_0=esk3_0|esk2_0=esk4_0|esk2_0=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70])])).
cnf(c_0_75, plain, (r1_tarski(k1_tarski(X1),k2_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_76, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|esk2_0=esk1_0|k1_tarski(X1)!=k1_tarski(esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_39]), c_0_72])).
cnf(c_0_77, plain, (r1_tarski(X1,k5_enumset1(X2,X2,X2,X2,X2,X2,X3))|X1!=k5_enumset1(X2,X2,X2,X2,X2,X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_23]), c_0_23]), c_0_24]), c_0_24])).
cnf(c_0_78, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|esk2_0=esk4_0|esk2_0=esk3_0|X1!=esk3_0), inference(ef,[status(thm)],[c_0_74])).
cnf(c_0_79, plain, (r1_tarski(k1_tarski(X1),k5_enumset1(X1,X1,X1,X1,X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_23]), c_0_24])).
cnf(c_0_80, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|esk2_0=esk1_0), inference(er,[status(thm)],[c_0_76])).
cnf(c_0_81, plain, (r1_tarski(X1,k2_tarski(X2,X3))|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_82, plain, (X1=X2|X1=X3|k1_tarski(X1)!=k5_enumset1(X2,X2,X2,X2,X2,X2,X3)), inference(spm,[status(thm)],[c_0_38, c_0_77])).
cnf(c_0_83, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|esk2_0=esk3_0|esk2_0=esk4_0), inference(er,[status(thm)],[c_0_78])).
cnf(c_0_84, plain, (k5_enumset1(X1,X1,X1,X1,X1,X1,X2)=k1_tarski(X1)|~r1_tarski(k5_enumset1(X1,X1,X1,X1,X1,X1,X2),k1_tarski(X1))), inference(spm,[status(thm)],[c_0_55, c_0_79])).
cnf(c_0_85, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|esk2_0=esk1_0|r1_tarski(k1_tarski(esk1_0),k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_86, plain, (r1_tarski(X1,k5_enumset1(X2,X2,X2,X2,X2,X2,X3))|X1!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_23]), c_0_24])).
cnf(c_0_87, negated_conjecture, (esk2_0=esk4_0|esk2_0=esk3_0|X1=esk2_0|X1=esk1_0|k1_tarski(X1)!=k1_tarski(esk3_0)), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_88, negated_conjecture, (esk1_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_89, negated_conjecture, (esk2_0=esk1_0|r1_tarski(k1_tarski(esk1_0),k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_64])]), c_0_86])).
cnf(c_0_90, negated_conjecture, (esk2_0=esk3_0|esk2_0=esk4_0), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_87]), c_0_88])).
cnf(c_0_91, negated_conjecture, (esk2_0=esk3_0|r1_tarski(k1_tarski(esk1_0),k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_46])).
cnf(c_0_92, negated_conjecture, (r1_tarski(k1_tarski(esk1_0),k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_91]), c_0_88])).
cnf(c_0_93, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_92]), c_0_46]), c_0_88]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 94
# Proof object clause steps            : 65
# Proof object formula steps           : 29
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 34
# Proof object simplifying inferences  : 38
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 24
# Processed clauses                    : 2403
# ...of these trivial                  : 22
# ...subsumed                          : 1455
# ...remaining for further processing  : 926
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 355
# Backward-rewritten                   : 18
# Generated clauses                    : 32118
# ...of the previous two non-trivial   : 28047
# Contextual simplify-reflections      : 50
# Paramodulations                      : 31952
# Factorizations                       : 108
# Equation resolutions                 : 58
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 528
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 509
# Current number of unprocessed clauses: 24656
# ...number of literals in the above   : 152564
# Current number of archived formulas  : 0
# Current number of archived clauses   : 398
# Clause-clause subsumption calls (NU) : 40504
# Rec. Clause-clause subsumption calls : 14495
# Non-unit clause-clause subsumptions  : 1677
# Unit Clause-clause subsumption calls : 554
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 15
# BW rewrite match successes           : 8
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1128144

# -------------------------------------------------
# User time                : 0.415 s
# System time              : 0.007 s
# Total time               : 0.421 s
# Maximum resident set size: 26188 pages
