# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C41_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/009783833dd54529b081dc9b0f2281ba.E_input', d3_tarski)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/009783833dd54529b081dc9b0f2281ba.E_input', d4_xboole_0)).
fof(d6_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:k2_wellord1(X1,X2)=k3_xboole_0(X1,k2_zfmisc_1(X2,X2))), file('data/MPTP/proofs/009783833dd54529b081dc9b0f2281ba.E_input', d6_wellord1)).
fof(t8_wellord2, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_wellord1(k1_wellord2(X2),X1)=k1_wellord2(X1)), file('data/MPTP/proofs/009783833dd54529b081dc9b0f2281ba.E_input', t8_wellord2)).
fof(dt_k1_wellord2, axiom, ![X1]:v1_relat_1(k1_wellord2(X1)), file('data/MPTP/proofs/009783833dd54529b081dc9b0f2281ba.E_input', dt_k1_wellord2)).
fof(t33_wellord2, conjecture, ![X1]:r1_tarski(k1_wellord2(X1),k2_zfmisc_1(X1,X1)), file('data/MPTP/proofs/009783833dd54529b081dc9b0f2281ba.E_input', t33_wellord2)).
fof(c_0_6, plain, ![X6, X7, X8, X9, X10]:((~r1_tarski(X6,X7)|(~r2_hidden(X8,X6)|r2_hidden(X8,X7)))&((r2_hidden(esk2_2(X9,X10),X9)|r1_tarski(X9,X10))&(~r2_hidden(esk2_2(X9,X10),X10)|r1_tarski(X9,X10)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_7, plain, ![X12, X13, X14, X15, X16, X17, X18, X19]:((((r2_hidden(X15,X12)|~r2_hidden(X15,X14)|X14!=k3_xboole_0(X12,X13))&(r2_hidden(X15,X13)|~r2_hidden(X15,X14)|X14!=k3_xboole_0(X12,X13)))&(~r2_hidden(X16,X12)|~r2_hidden(X16,X13)|r2_hidden(X16,X14)|X14!=k3_xboole_0(X12,X13)))&((~r2_hidden(esk3_3(X17,X18,X19),X19)|(~r2_hidden(esk3_3(X17,X18,X19),X17)|~r2_hidden(esk3_3(X17,X18,X19),X18))|X19=k3_xboole_0(X17,X18))&((r2_hidden(esk3_3(X17,X18,X19),X17)|r2_hidden(esk3_3(X17,X18,X19),X19)|X19=k3_xboole_0(X17,X18))&(r2_hidden(esk3_3(X17,X18,X19),X18)|r2_hidden(esk3_3(X17,X18,X19),X19)|X19=k3_xboole_0(X17,X18))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_8, plain, ![X24, X25]:(~v1_relat_1(X24)|k2_wellord1(X24,X25)=k3_xboole_0(X24,k2_zfmisc_1(X25,X25))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_wellord1])])])).
fof(c_0_9, plain, ![X22, X23]:(~r1_tarski(X22,X23)|k2_wellord1(k1_wellord2(X23),X22)=k1_wellord2(X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_wellord2])])).
cnf(c_0_10, plain, (r1_tarski(X1,X2)|~r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, plain, (k2_wellord1(X1,X2)=k3_xboole_0(X1,k2_zfmisc_1(X2,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, plain, (k2_wellord1(k1_wellord2(X2),X1)=k1_wellord2(X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
fof(c_0_16, plain, ![X21]:v1_relat_1(k1_wellord2(X21)), inference(variable_rename,[status(thm)],[dt_k1_wellord2])).
cnf(c_0_17, plain, (r2_hidden(X1,k2_zfmisc_1(X2,X2))|X3!=k2_wellord1(X4,X2)|~v1_relat_1(X4)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_18, plain, (k2_wellord1(k1_wellord2(X1),X1)=k1_wellord2(X1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_19, plain, (v1_relat_1(k1_wellord2(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, (r2_hidden(X1,k2_zfmisc_1(X2,X2))|X3!=k1_wellord2(X2)|~r2_hidden(X1,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])])).
fof(c_0_21, negated_conjecture, ~(![X1]:r1_tarski(k1_wellord2(X1),k2_zfmisc_1(X1,X1))), inference(assume_negation,[status(cth)],[t33_wellord2])).
cnf(c_0_22, plain, (r2_hidden(X1,k2_zfmisc_1(X2,X2))|~r2_hidden(X1,k1_wellord2(X2))), inference(er,[status(thm)],[c_0_20])).
fof(c_0_23, negated_conjecture, ~r1_tarski(k1_wellord2(esk1_0),k2_zfmisc_1(esk1_0,esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).
cnf(c_0_24, plain, (r2_hidden(esk2_2(k1_wellord2(X1),X2),k2_zfmisc_1(X1,X1))|r1_tarski(k1_wellord2(X1),X2)), inference(spm,[status(thm)],[c_0_22, c_0_11])).
cnf(c_0_25, negated_conjecture, (~r1_tarski(k1_wellord2(esk1_0),k2_zfmisc_1(esk1_0,esk1_0))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_26, plain, (r1_tarski(k1_wellord2(X1),k2_zfmisc_1(X1,X1))), inference(spm,[status(thm)],[c_0_10, c_0_24])).
cnf(c_0_27, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_26])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 28
# Proof object clause steps            : 15
# Proof object formula steps           : 13
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 7
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 4
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 13
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 13
# Processed clauses                    : 37
# ...of these trivial                  : 2
# ...subsumed                          : 3
# ...remaining for further processing  : 32
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 85
# ...of the previous two non-trivial   : 71
# Contextual simplify-reflections      : 1
# Paramodulations                      : 71
# Factorizations                       : 4
# Equation resolutions                 : 10
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 31
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 26
# Current number of unprocessed clauses: 47
# ...number of literals in the above   : 155
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1
# Clause-clause subsumption calls (NU) : 85
# Rec. Clause-clause subsumption calls : 50
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 11
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 13
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1794

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.003 s
# Total time               : 0.009 s
# Maximum resident set size: 28852 pages
