# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t23_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(m1_pre_topc(X2,X3)=>((r1_tsep_1(X2,X4)&r1_tsep_1(X4,X2))|(~(r1_tsep_1(X3,X4))&~(r1_tsep_1(X4,X3))))))))), file('data/MPTP/proofs/00bd17b4447d4f4189c68354fbce0407.E_input', t23_tmap_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/00bd17b4447d4f4189c68354fbce0407.E_input', d10_xboole_0)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/00bd17b4447d4f4189c68354fbce0407.E_input', dt_m1_pre_topc)).
fof(t4_tsep_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X1)=>(r1_tarski(u1_struct_0(X2),u1_struct_0(X3))<=>m1_pre_topc(X2,X3))))), file('data/MPTP/proofs/00bd17b4447d4f4189c68354fbce0407.E_input', t4_tsep_1)).
fof(dt_k1_tsep_1, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))&(~(v2_struct_0(X3))&m1_pre_topc(X3,X1))))=>(~(v2_struct_0(k1_tsep_1(X1,X2,X3)))&(v1_pre_topc(k1_tsep_1(X1,X2,X3))&m1_pre_topc(k1_tsep_1(X1,X2,X3),X1)))), file('data/MPTP/proofs/00bd17b4447d4f4189c68354fbce0407.E_input', dt_k1_tsep_1)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/00bd17b4447d4f4189c68354fbce0407.E_input', dt_l1_pre_topc)).
fof(d2_tsep_1, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&(v1_pre_topc(X4)&m1_pre_topc(X4,X1)))=>(X4=k1_tsep_1(X1,X2,X3)<=>u1_struct_0(X4)=k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))))), file('data/MPTP/proofs/00bd17b4447d4f4189c68354fbce0407.E_input', d2_tsep_1)).
fof(d9_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>(m1_pre_topc(X2,X1)<=>(r1_tarski(k2_struct_0(X2),k2_struct_0(X1))&![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>(r2_hidden(X3,u1_pre_topc(X2))<=>?[X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))&(r2_hidden(X4,u1_pre_topc(X1))&X3=k9_subset_1(u1_struct_0(X2),X4,k2_struct_0(X2)))))))))), file('data/MPTP/proofs/00bd17b4447d4f4189c68354fbce0407.E_input', d9_pre_topc)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/00bd17b4447d4f4189c68354fbce0407.E_input', d3_struct_0)).
fof(t11_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k2_xboole_0(X1,X2),X3)=>r1_tarski(X1,X3)), file('data/MPTP/proofs/00bd17b4447d4f4189c68354fbce0407.E_input', t11_xboole_1)).
fof(t70_xboole_1, axiom, ![X1, X2, X3]:(~((~(r1_xboole_0(X1,k2_xboole_0(X2,X3)))&(r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3))))&~((~((r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3)))&r1_xboole_0(X1,k2_xboole_0(X2,X3))))), file('data/MPTP/proofs/00bd17b4447d4f4189c68354fbce0407.E_input', t70_xboole_1)).
fof(symmetry_r1_tsep_1, axiom, ![X1, X2]:((l1_struct_0(X1)&l1_struct_0(X2))=>(r1_tsep_1(X1,X2)=>r1_tsep_1(X2,X1))), file('data/MPTP/proofs/00bd17b4447d4f4189c68354fbce0407.E_input', symmetry_r1_tsep_1)).
fof(d3_tsep_1, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(l1_struct_0(X2)=>(r1_tsep_1(X1,X2)<=>r1_xboole_0(u1_struct_0(X1),u1_struct_0(X2))))), file('data/MPTP/proofs/00bd17b4447d4f4189c68354fbce0407.E_input', d3_tsep_1)).
fof(c_0_13, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(m1_pre_topc(X2,X3)=>((r1_tsep_1(X2,X4)&r1_tsep_1(X4,X2))|(~(r1_tsep_1(X3,X4))&~(r1_tsep_1(X4,X3)))))))))), inference(assume_negation,[status(cth)],[t23_tmap_1])).
fof(c_0_14, plain, ![X32, X33]:(((r1_tarski(X32,X33)|X32!=X33)&(r1_tarski(X33,X32)|X32!=X33))&(~r1_tarski(X32,X33)|~r1_tarski(X33,X32)|X32=X33)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_15, plain, ![X35, X36]:(~l1_pre_topc(X35)|(~m1_pre_topc(X36,X35)|l1_pre_topc(X36))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
fof(c_0_16, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk1_0))&((~v2_struct_0(esk4_0)&m1_pre_topc(esk4_0,esk1_0))&(m1_pre_topc(esk2_0,esk3_0)&((~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk4_0,esk2_0))&(r1_tsep_1(esk3_0,esk4_0)|r1_tsep_1(esk4_0,esk3_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_13])])])])).
fof(c_0_17, plain, ![X21, X22, X23]:((~r1_tarski(u1_struct_0(X22),u1_struct_0(X23))|m1_pre_topc(X22,X23)|~m1_pre_topc(X23,X21)|~m1_pre_topc(X22,X21)|(~v2_pre_topc(X21)|~l1_pre_topc(X21)))&(~m1_pre_topc(X22,X23)|r1_tarski(u1_struct_0(X22),u1_struct_0(X23))|~m1_pre_topc(X23,X21)|~m1_pre_topc(X22,X21)|(~v2_pre_topc(X21)|~l1_pre_topc(X21)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_tsep_1])])])])).
cnf(c_0_18, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_19, plain, ![X37, X38, X39]:((~v2_struct_0(k1_tsep_1(X37,X38,X39))|(v2_struct_0(X37)|~l1_pre_topc(X37)|(v2_struct_0(X38)|~m1_pre_topc(X38,X37)|(v2_struct_0(X39)|~m1_pre_topc(X39,X37)))))&((v1_pre_topc(k1_tsep_1(X37,X38,X39))|(v2_struct_0(X37)|~l1_pre_topc(X37)|(v2_struct_0(X38)|~m1_pre_topc(X38,X37)|(v2_struct_0(X39)|~m1_pre_topc(X39,X37)))))&(m1_pre_topc(k1_tsep_1(X37,X38,X39),X37)|(v2_struct_0(X37)|~l1_pre_topc(X37)|(v2_struct_0(X38)|~m1_pre_topc(X38,X37)|(v2_struct_0(X39)|~m1_pre_topc(X39,X37))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_tsep_1])])])])).
cnf(c_0_20, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (m1_pre_topc(X1,X2)|~r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (m1_pre_topc(k1_tsep_1(X1,X2,X3),X1)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, negated_conjecture, (m1_pre_topc(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_27, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_29, negated_conjecture, (l1_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])])).
cnf(c_0_30, plain, (m1_pre_topc(X1,X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_31, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_32, plain, ![X9]:(~l1_pre_topc(X9)|l1_struct_0(X9)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_33, negated_conjecture, (m1_pre_topc(k1_tsep_1(esk3_0,X1,esk2_0),esk3_0)|v2_struct_0(X1)|~m1_pre_topc(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_28]), c_0_29])])).
cnf(c_0_34, negated_conjecture, (m1_pre_topc(esk3_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_21]), c_0_22]), c_0_31])])).
fof(c_0_35, plain, ![X13, X14, X15, X16]:((X16!=k1_tsep_1(X13,X14,X15)|u1_struct_0(X16)=k2_xboole_0(u1_struct_0(X14),u1_struct_0(X15))|(v2_struct_0(X16)|(~v1_pre_topc(X16)|~m1_pre_topc(X16,X13)))|(v2_struct_0(X15)|~m1_pre_topc(X15,X13))|(v2_struct_0(X14)|~m1_pre_topc(X14,X13))|(v2_struct_0(X13)|~l1_pre_topc(X13)))&(u1_struct_0(X16)!=k2_xboole_0(u1_struct_0(X14),u1_struct_0(X15))|X16=k1_tsep_1(X13,X14,X15)|(v2_struct_0(X16)|(~v1_pre_topc(X16)|~m1_pre_topc(X16,X13)))|(v2_struct_0(X15)|~m1_pre_topc(X15,X13))|(v2_struct_0(X14)|~m1_pre_topc(X14,X13))|(v2_struct_0(X13)|~l1_pre_topc(X13)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_tsep_1])])])])])).
fof(c_0_36, plain, ![X24, X25, X26, X28, X30]:(((r1_tarski(k2_struct_0(X25),k2_struct_0(X24))|~m1_pre_topc(X25,X24)|~l1_pre_topc(X25)|~l1_pre_topc(X24))&(((m1_subset_1(esk5_3(X24,X25,X26),k1_zfmisc_1(u1_struct_0(X24)))|~r2_hidden(X26,u1_pre_topc(X25))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|~m1_pre_topc(X25,X24)|~l1_pre_topc(X25)|~l1_pre_topc(X24))&((r2_hidden(esk5_3(X24,X25,X26),u1_pre_topc(X24))|~r2_hidden(X26,u1_pre_topc(X25))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|~m1_pre_topc(X25,X24)|~l1_pre_topc(X25)|~l1_pre_topc(X24))&(X26=k9_subset_1(u1_struct_0(X25),esk5_3(X24,X25,X26),k2_struct_0(X25))|~r2_hidden(X26,u1_pre_topc(X25))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|~m1_pre_topc(X25,X24)|~l1_pre_topc(X25)|~l1_pre_topc(X24))))&(~m1_subset_1(X28,k1_zfmisc_1(u1_struct_0(X24)))|(~r2_hidden(X28,u1_pre_topc(X24))|X26!=k9_subset_1(u1_struct_0(X25),X28,k2_struct_0(X25)))|r2_hidden(X26,u1_pre_topc(X25))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|~m1_pre_topc(X25,X24)|~l1_pre_topc(X25)|~l1_pre_topc(X24))))&((m1_subset_1(esk6_2(X24,X25),k1_zfmisc_1(u1_struct_0(X25)))|~r1_tarski(k2_struct_0(X25),k2_struct_0(X24))|m1_pre_topc(X25,X24)|~l1_pre_topc(X25)|~l1_pre_topc(X24))&((~r2_hidden(esk6_2(X24,X25),u1_pre_topc(X25))|(~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X24)))|(~r2_hidden(X30,u1_pre_topc(X24))|esk6_2(X24,X25)!=k9_subset_1(u1_struct_0(X25),X30,k2_struct_0(X25))))|~r1_tarski(k2_struct_0(X25),k2_struct_0(X24))|m1_pre_topc(X25,X24)|~l1_pre_topc(X25)|~l1_pre_topc(X24))&((m1_subset_1(esk7_2(X24,X25),k1_zfmisc_1(u1_struct_0(X24)))|r2_hidden(esk6_2(X24,X25),u1_pre_topc(X25))|~r1_tarski(k2_struct_0(X25),k2_struct_0(X24))|m1_pre_topc(X25,X24)|~l1_pre_topc(X25)|~l1_pre_topc(X24))&((r2_hidden(esk7_2(X24,X25),u1_pre_topc(X24))|r2_hidden(esk6_2(X24,X25),u1_pre_topc(X25))|~r1_tarski(k2_struct_0(X25),k2_struct_0(X24))|m1_pre_topc(X25,X24)|~l1_pre_topc(X25)|~l1_pre_topc(X24))&(esk6_2(X24,X25)=k9_subset_1(u1_struct_0(X25),esk7_2(X24,X25),k2_struct_0(X25))|r2_hidden(esk6_2(X24,X25),u1_pre_topc(X25))|~r1_tarski(k2_struct_0(X25),k2_struct_0(X24))|m1_pre_topc(X25,X24)|~l1_pre_topc(X25)|~l1_pre_topc(X24))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_pre_topc])])])])])).
fof(c_0_37, plain, ![X34]:(~l1_struct_0(X34)|k2_struct_0(X34)=u1_struct_0(X34)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_38, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39, negated_conjecture, (m1_pre_topc(k1_tsep_1(esk3_0,esk3_0,esk2_0),esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_28])).
cnf(c_0_40, plain, (u1_struct_0(X1)=k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4))|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|X1!=k1_tsep_1(X2,X3,X4)|~v1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41, plain, (v1_pre_topc(k1_tsep_1(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_42, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~v2_struct_0(k1_tsep_1(X1,X2,X3))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_43, plain, ![X40, X41, X42]:(~r1_tarski(k2_xboole_0(X40,X41),X42)|r1_tarski(X40,X42)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_xboole_1])])).
cnf(c_0_44, plain, (r1_tarski(k2_struct_0(X1),k2_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X1)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_45, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_46, negated_conjecture, (l1_struct_0(esk3_0)), inference(spm,[status(thm)],[c_0_38, c_0_29])).
cnf(c_0_47, negated_conjecture, (l1_pre_topc(k1_tsep_1(esk3_0,esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_39]), c_0_29])])).
cnf(c_0_48, plain, (u1_struct_0(k1_tsep_1(X1,X2,X3))=k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_40]), c_0_25]), c_0_41]), c_0_42])).
cnf(c_0_49, plain, (r1_tarski(X1,X3)|~r1_tarski(k2_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_50, plain, (r1_tarski(k2_struct_0(X1),k2_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[c_0_44, c_0_20])).
cnf(c_0_51, negated_conjecture, (k2_struct_0(esk3_0)=u1_struct_0(esk3_0)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_52, negated_conjecture, (l1_struct_0(k1_tsep_1(esk3_0,esk3_0,esk2_0))), inference(spm,[status(thm)],[c_0_38, c_0_47])).
cnf(c_0_53, negated_conjecture, (m1_pre_topc(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_54, plain, ![X10, X11, X12]:((r1_xboole_0(X10,k2_xboole_0(X11,X12))|(~r1_xboole_0(X10,X11)|~r1_xboole_0(X10,X12)))&((r1_xboole_0(X10,X11)|~r1_xboole_0(X10,k2_xboole_0(X11,X12)))&(r1_xboole_0(X10,X12)|~r1_xboole_0(X10,k2_xboole_0(X11,X12))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t70_xboole_1])])])])).
cnf(c_0_55, negated_conjecture, (k2_xboole_0(u1_struct_0(X1),u1_struct_0(esk2_0))=u1_struct_0(k1_tsep_1(esk3_0,X1,esk2_0))|v2_struct_0(X1)|~m1_pre_topc(X1,esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_26]), c_0_29])]), c_0_28]), c_0_27])).
cnf(c_0_56, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_57, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_49, c_0_24])).
cnf(c_0_58, negated_conjecture, (r1_tarski(k2_struct_0(k1_tsep_1(esk3_0,esk3_0,esk2_0)),u1_struct_0(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_39]), c_0_51]), c_0_29])])).
cnf(c_0_59, negated_conjecture, (k2_struct_0(k1_tsep_1(esk3_0,esk3_0,esk2_0))=u1_struct_0(k1_tsep_1(esk3_0,esk3_0,esk2_0))), inference(spm,[status(thm)],[c_0_45, c_0_52])).
fof(c_0_60, plain, ![X19, X20]:(~l1_struct_0(X19)|~l1_struct_0(X20)|(~r1_tsep_1(X19,X20)|r1_tsep_1(X20,X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_tsep_1])])).
cnf(c_0_61, negated_conjecture, (l1_pre_topc(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_53]), c_0_22])])).
cnf(c_0_62, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X1,k2_xboole_0(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_63, negated_conjecture, (k2_xboole_0(u1_struct_0(esk3_0),u1_struct_0(esk2_0))=u1_struct_0(k1_tsep_1(esk3_0,esk3_0,esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_34]), c_0_28])).
cnf(c_0_64, plain, (k2_xboole_0(X1,X2)=X1|~r1_tarski(k2_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_65, negated_conjecture, (r1_tarski(u1_struct_0(k1_tsep_1(esk3_0,esk3_0,esk2_0)),u1_struct_0(esk3_0))), inference(rw,[status(thm)],[c_0_58, c_0_59])).
fof(c_0_66, plain, ![X17, X18]:((~r1_tsep_1(X17,X18)|r1_xboole_0(u1_struct_0(X17),u1_struct_0(X18))|~l1_struct_0(X18)|~l1_struct_0(X17))&(~r1_xboole_0(u1_struct_0(X17),u1_struct_0(X18))|r1_tsep_1(X17,X18)|~l1_struct_0(X18)|~l1_struct_0(X17))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tsep_1])])])])).
cnf(c_0_67, plain, (r1_tsep_1(X2,X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~r1_tsep_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_68, negated_conjecture, (r1_tsep_1(esk3_0,esk4_0)|r1_tsep_1(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_69, negated_conjecture, (l1_struct_0(esk4_0)), inference(spm,[status(thm)],[c_0_38, c_0_61])).
cnf(c_0_70, negated_conjecture, (r1_xboole_0(X1,u1_struct_0(esk2_0))|~r1_xboole_0(X1,u1_struct_0(k1_tsep_1(esk3_0,esk3_0,esk2_0)))), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_71, negated_conjecture, (u1_struct_0(k1_tsep_1(esk3_0,esk3_0,esk2_0))=u1_struct_0(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_63]), c_0_65])])).
cnf(c_0_72, plain, (r1_xboole_0(u1_struct_0(X1),u1_struct_0(X2))|~r1_tsep_1(X1,X2)|~l1_struct_0(X2)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_73, negated_conjecture, (r1_tsep_1(esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_69]), c_0_46])])).
cnf(c_0_74, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_75, negated_conjecture, (r1_xboole_0(X1,u1_struct_0(esk2_0))|~r1_xboole_0(X1,u1_struct_0(esk3_0))), inference(rw,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_76, negated_conjecture, (r1_xboole_0(u1_struct_0(esk4_0),u1_struct_0(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_46]), c_0_69])])).
cnf(c_0_77, negated_conjecture, (l1_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_74]), c_0_22])])).
cnf(c_0_78, plain, (r1_tsep_1(X1,X2)|~r1_xboole_0(u1_struct_0(X1),u1_struct_0(X2))|~l1_struct_0(X2)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_79, negated_conjecture, (r1_xboole_0(u1_struct_0(esk4_0),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_80, negated_conjecture, (l1_struct_0(esk2_0)), inference(spm,[status(thm)],[c_0_38, c_0_77])).
cnf(c_0_81, negated_conjecture, (~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_82, negated_conjecture, (r1_tsep_1(esk4_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80]), c_0_69])])).
cnf(c_0_83, negated_conjecture, (~r1_tsep_1(esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_82])])).
cnf(c_0_84, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_82]), c_0_80]), c_0_69])]), c_0_83]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 85
# Proof object clause steps            : 58
# Proof object formula steps           : 27
# Proof object conjectures             : 39
# Proof object clause conjectures      : 36
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 48
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 42
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 42
# Processed clauses                    : 450
# ...of these trivial                  : 12
# ...subsumed                          : 7
# ...remaining for further processing  : 430
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 38
# Generated clauses                    : 1928
# ...of the previous two non-trivial   : 1750
# Contextual simplify-reflections      : 8
# Paramodulations                      : 1923
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 349
#    Positive orientable unit clauses  : 162
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 182
# Current number of unprocessed clauses: 1378
# ...number of literals in the above   : 5056
# Current number of archived formulas  : 0
# Current number of archived clauses   : 79
# Clause-clause subsumption calls (NU) : 3814
# Rec. Clause-clause subsumption calls : 1979
# Non-unit clause-clause subsumptions  : 15
# Unit Clause-clause subsumption calls : 842
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 153
# BW rewrite match successes           : 18
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 61702

# -------------------------------------------------
# User time                : 0.025 s
# System time              : 0.007 s
# Total time               : 0.032 s
# Maximum resident set size: 28852 pages
