# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t156_funct_1, conjecture, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>![X4]:((v1_relat_1(X4)&v1_funct_1(X4))=>((X1=k2_relat_1(X2)&(k1_relat_1(X3)=X1&(k1_relat_1(X4)=X1&k5_relat_1(X2,X3)=k5_relat_1(X2,X4))))=>X3=X4)))), file('data/MPTP/proofs/00d19e060fe2404eb6f7aa631a649e7d.E_input', t156_funct_1)).
fof(d3_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(r1_tarski(X1,X2)<=>![X3, X4]:(r2_hidden(k4_tarski(X3,X4),X1)=>r2_hidden(k4_tarski(X3,X4),X2)))), file('data/MPTP/proofs/00d19e060fe2404eb6f7aa631a649e7d.E_input', d3_relat_1)).
fof(t20_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)=>(r2_hidden(X1,k1_relat_1(X3))&r2_hidden(X2,k2_relat_1(X3))))), file('data/MPTP/proofs/00d19e060fe2404eb6f7aa631a649e7d.E_input', t20_relat_1)).
fof(d5_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4))))), file('data/MPTP/proofs/00d19e060fe2404eb6f7aa631a649e7d.E_input', d5_funct_1)).
fof(t23_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k1_relat_1(X2))=>k1_funct_1(k5_relat_1(X2,X3),X1)=k1_funct_1(X3,k1_funct_1(X2,X1))))), file('data/MPTP/proofs/00d19e060fe2404eb6f7aa631a649e7d.E_input', t23_funct_1)).
fof(t8_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>(r2_hidden(X1,k1_relat_1(X3))&X2=k1_funct_1(X3,X1)))), file('data/MPTP/proofs/00d19e060fe2404eb6f7aa631a649e7d.E_input', t8_funct_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/00d19e060fe2404eb6f7aa631a649e7d.E_input', d10_xboole_0)).
fof(c_0_7, negated_conjecture, ~(![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>![X4]:((v1_relat_1(X4)&v1_funct_1(X4))=>((X1=k2_relat_1(X2)&(k1_relat_1(X3)=X1&(k1_relat_1(X4)=X1&k5_relat_1(X2,X3)=k5_relat_1(X2,X4))))=>X3=X4))))), inference(assume_negation,[status(cth)],[t156_funct_1])).
fof(c_0_8, plain, ![X14, X15, X16, X17, X18]:((~r1_tarski(X14,X15)|(~r2_hidden(k4_tarski(X16,X17),X14)|r2_hidden(k4_tarski(X16,X17),X15))|~v1_relat_1(X14))&((r2_hidden(k4_tarski(esk5_2(X14,X18),esk6_2(X14,X18)),X14)|r1_tarski(X14,X18)|~v1_relat_1(X14))&(~r2_hidden(k4_tarski(esk5_2(X14,X18),esk6_2(X14,X18)),X18)|r1_tarski(X14,X18)|~v1_relat_1(X14)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_relat_1])])])])])])).
fof(c_0_9, negated_conjecture, ((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&((v1_relat_1(esk3_0)&v1_funct_1(esk3_0))&((v1_relat_1(esk4_0)&v1_funct_1(esk4_0))&((esk1_0=k2_relat_1(esk2_0)&(k1_relat_1(esk3_0)=esk1_0&(k1_relat_1(esk4_0)=esk1_0&k5_relat_1(esk2_0,esk3_0)=k5_relat_1(esk2_0,esk4_0))))&esk3_0!=esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
fof(c_0_10, plain, ![X24, X25, X26]:((r2_hidden(X24,k1_relat_1(X26))|~r2_hidden(k4_tarski(X24,X25),X26)|~v1_relat_1(X26))&(r2_hidden(X25,k2_relat_1(X26))|~r2_hidden(k4_tarski(X24,X25),X26)|~v1_relat_1(X26))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t20_relat_1])])])).
cnf(c_0_11, plain, (r2_hidden(k4_tarski(esk5_2(X1,X2),esk6_2(X1,X2)),X1)|r1_tarski(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, negated_conjecture, (v1_relat_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_13, plain, ![X27, X28, X29, X31, X32, X33, X35]:((((r2_hidden(esk7_3(X27,X28,X29),k1_relat_1(X27))|~r2_hidden(X29,X28)|X28!=k2_relat_1(X27)|(~v1_relat_1(X27)|~v1_funct_1(X27)))&(X29=k1_funct_1(X27,esk7_3(X27,X28,X29))|~r2_hidden(X29,X28)|X28!=k2_relat_1(X27)|(~v1_relat_1(X27)|~v1_funct_1(X27))))&(~r2_hidden(X32,k1_relat_1(X27))|X31!=k1_funct_1(X27,X32)|r2_hidden(X31,X28)|X28!=k2_relat_1(X27)|(~v1_relat_1(X27)|~v1_funct_1(X27))))&((~r2_hidden(esk8_2(X27,X33),X33)|(~r2_hidden(X35,k1_relat_1(X27))|esk8_2(X27,X33)!=k1_funct_1(X27,X35))|X33=k2_relat_1(X27)|(~v1_relat_1(X27)|~v1_funct_1(X27)))&((r2_hidden(esk9_2(X27,X33),k1_relat_1(X27))|r2_hidden(esk8_2(X27,X33),X33)|X33=k2_relat_1(X27)|(~v1_relat_1(X27)|~v1_funct_1(X27)))&(esk8_2(X27,X33)=k1_funct_1(X27,esk9_2(X27,X33))|r2_hidden(esk8_2(X27,X33),X33)|X33=k2_relat_1(X27)|(~v1_relat_1(X27)|~v1_funct_1(X27)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_funct_1])])])])])])).
cnf(c_0_14, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (r2_hidden(k4_tarski(esk5_2(esk4_0,X1),esk6_2(esk4_0,X1)),esk4_0)|r1_tarski(esk4_0,X1)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16, negated_conjecture, (k1_relat_1(esk4_0)=esk1_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, plain, (r2_hidden(esk7_3(X1,X2,X3),k1_relat_1(X1))|~r2_hidden(X3,X2)|X2!=k2_relat_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, negated_conjecture, (r2_hidden(esk5_2(esk4_0,X1),esk1_0)|r1_tarski(esk4_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_12])])).
fof(c_0_19, plain, ![X11, X12, X13]:(~v1_relat_1(X12)|~v1_funct_1(X12)|(~v1_relat_1(X13)|~v1_funct_1(X13)|(~r2_hidden(X11,k1_relat_1(X12))|k1_funct_1(k5_relat_1(X12,X13),X11)=k1_funct_1(X13,k1_funct_1(X12,X11))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_funct_1])])])).
cnf(c_0_20, negated_conjecture, (r2_hidden(esk7_3(X1,esk1_0,esk5_2(esk4_0,X2)),k1_relat_1(X1))|r1_tarski(esk4_0,X2)|k2_relat_1(X1)!=esk1_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_21, negated_conjecture, (esk1_0=k2_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_22, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_23, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_24, plain, (k1_funct_1(k5_relat_1(X1,X2),X3)=k1_funct_1(X2,k1_funct_1(X1,X3))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~r2_hidden(X3,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, negated_conjecture, (r2_hidden(esk7_3(esk2_0,esk1_0,esk5_2(esk4_0,X1)),k1_relat_1(esk2_0))|r1_tarski(esk4_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23])])).
fof(c_0_26, plain, ![X21, X22, X23]:(((r2_hidden(X21,k1_relat_1(X23))|~r2_hidden(k4_tarski(X21,X22),X23)|(~v1_relat_1(X23)|~v1_funct_1(X23)))&(X22=k1_funct_1(X23,X21)|~r2_hidden(k4_tarski(X21,X22),X23)|(~v1_relat_1(X23)|~v1_funct_1(X23))))&(~r2_hidden(X21,k1_relat_1(X23))|X22!=k1_funct_1(X23,X21)|r2_hidden(k4_tarski(X21,X22),X23)|(~v1_relat_1(X23)|~v1_funct_1(X23)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_funct_1])])])).
cnf(c_0_27, negated_conjecture, (k1_funct_1(k5_relat_1(esk2_0,X1),esk7_3(esk2_0,esk1_0,esk5_2(esk4_0,X2)))=k1_funct_1(X1,k1_funct_1(esk2_0,esk7_3(esk2_0,esk1_0,esk5_2(esk4_0,X2))))|r1_tarski(esk4_0,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_22]), c_0_23])])).
cnf(c_0_28, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_29, negated_conjecture, (k5_relat_1(esk2_0,esk3_0)=k5_relat_1(esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_30, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_31, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_32, plain, (X1=k1_funct_1(X2,esk7_3(X2,X3,X1))|~r2_hidden(X1,X3)|X3!=k2_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_33, plain, (r2_hidden(k4_tarski(X1,X3),X2)|~r2_hidden(X1,k1_relat_1(X2))|X3!=k1_funct_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, negated_conjecture, (k1_relat_1(esk3_0)=esk1_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_35, negated_conjecture, (k1_funct_1(k5_relat_1(esk2_0,esk3_0),esk7_3(esk2_0,esk1_0,esk5_2(esk4_0,X1)))=k1_funct_1(esk4_0,k1_funct_1(esk2_0,esk7_3(esk2_0,esk1_0,esk5_2(esk4_0,X1))))|r1_tarski(esk4_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_12])])).
cnf(c_0_36, negated_conjecture, (k1_funct_1(k5_relat_1(esk2_0,esk3_0),esk7_3(esk2_0,esk1_0,esk5_2(esk4_0,X1)))=k1_funct_1(esk3_0,k1_funct_1(esk2_0,esk7_3(esk2_0,esk1_0,esk5_2(esk4_0,X1))))|r1_tarski(esk4_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_30]), c_0_31])])).
cnf(c_0_37, negated_conjecture, (k1_funct_1(X1,esk7_3(X1,esk1_0,esk5_2(esk4_0,X2)))=esk5_2(esk4_0,X2)|r1_tarski(esk4_0,X2)|k2_relat_1(X1)!=esk1_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_32, c_0_18])).
cnf(c_0_38, negated_conjecture, (r2_hidden(k4_tarski(X1,X2),esk3_0)|X2!=k1_funct_1(esk3_0,X1)|~r2_hidden(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_30]), c_0_31])])).
cnf(c_0_39, plain, (X1=k1_funct_1(X2,X3)|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_40, negated_conjecture, (k1_funct_1(esk4_0,k1_funct_1(esk2_0,esk7_3(esk2_0,esk1_0,esk5_2(esk4_0,X1))))=k1_funct_1(esk3_0,k1_funct_1(esk2_0,esk7_3(esk2_0,esk1_0,esk5_2(esk4_0,X1))))|r1_tarski(esk4_0,X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_41, negated_conjecture, (k1_funct_1(esk2_0,esk7_3(esk2_0,esk1_0,esk5_2(esk4_0,X1)))=esk5_2(esk4_0,X1)|r1_tarski(esk4_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_21]), c_0_22]), c_0_23])])).
cnf(c_0_42, negated_conjecture, (r2_hidden(k4_tarski(X1,X2),esk4_0)|X2!=k1_funct_1(esk4_0,X1)|~r2_hidden(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_16]), c_0_28]), c_0_12])])).
cnf(c_0_43, negated_conjecture, (r2_hidden(k4_tarski(esk5_2(esk3_0,X1),esk6_2(esk3_0,X1)),esk3_0)|r1_tarski(esk3_0,X1)), inference(spm,[status(thm)],[c_0_11, c_0_31])).
cnf(c_0_44, negated_conjecture, (r2_hidden(k4_tarski(X1,k1_funct_1(esk3_0,X1)),esk3_0)|~r2_hidden(X1,esk1_0)), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_45, negated_conjecture, (k1_funct_1(esk4_0,esk5_2(esk4_0,X1))=esk6_2(esk4_0,X1)|r1_tarski(esk4_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_15]), c_0_28]), c_0_12])])).
cnf(c_0_46, negated_conjecture, (k1_funct_1(esk4_0,esk5_2(esk4_0,X1))=k1_funct_1(esk3_0,esk5_2(esk4_0,X1))|r1_tarski(esk4_0,X1)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_47, negated_conjecture, (r2_hidden(k4_tarski(X1,k1_funct_1(esk4_0,X1)),esk4_0)|~r2_hidden(X1,esk1_0)), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_48, negated_conjecture, (r2_hidden(esk5_2(esk3_0,X1),esk1_0)|r1_tarski(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_43]), c_0_34]), c_0_31])])).
cnf(c_0_49, negated_conjecture, (r2_hidden(k4_tarski(esk5_2(esk4_0,X1),k1_funct_1(esk3_0,esk5_2(esk4_0,X1))),esk3_0)|r1_tarski(esk4_0,X1)), inference(spm,[status(thm)],[c_0_44, c_0_18])).
cnf(c_0_50, negated_conjecture, (k1_funct_1(esk3_0,esk5_2(esk4_0,X1))=esk6_2(esk4_0,X1)|r1_tarski(esk4_0,X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_51, plain, (r2_hidden(k4_tarski(X3,X4),X2)|~r1_tarski(X1,X2)|~r2_hidden(k4_tarski(X3,X4),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_52, negated_conjecture, (r2_hidden(k4_tarski(esk5_2(esk3_0,X1),k1_funct_1(esk4_0,esk5_2(esk3_0,X1))),esk4_0)|r1_tarski(esk3_0,X1)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_53, plain, (r1_tarski(X1,X2)|~r2_hidden(k4_tarski(esk5_2(X1,X2),esk6_2(X1,X2)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_54, negated_conjecture, (r2_hidden(k4_tarski(esk5_2(esk4_0,X1),esk6_2(esk4_0,X1)),esk3_0)|r1_tarski(esk4_0,X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_55, negated_conjecture, (r2_hidden(k4_tarski(esk5_2(esk3_0,X1),k1_funct_1(esk4_0,esk5_2(esk3_0,X1))),X2)|r1_tarski(esk3_0,X1)|~r1_tarski(esk4_0,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_12])])).
cnf(c_0_56, negated_conjecture, (r1_tarski(esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_12])])).
cnf(c_0_57, negated_conjecture, (r2_hidden(k4_tarski(esk5_2(esk3_0,X1),k1_funct_1(esk4_0,esk5_2(esk3_0,X1))),esk3_0)|r1_tarski(esk3_0,X1)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_58, negated_conjecture, (k1_funct_1(esk4_0,esk5_2(esk3_0,X1))=k1_funct_1(esk3_0,esk5_2(esk3_0,X1))|r1_tarski(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_57]), c_0_30]), c_0_31])])).
cnf(c_0_59, negated_conjecture, (r2_hidden(k4_tarski(esk5_2(esk3_0,X1),X2),esk4_0)|r1_tarski(esk3_0,X1)|X2!=k1_funct_1(esk3_0,esk5_2(esk3_0,X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_58]), c_0_48])).
cnf(c_0_60, negated_conjecture, (k1_funct_1(esk3_0,esk5_2(esk3_0,X1))=esk6_2(esk3_0,X1)|r1_tarski(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_43]), c_0_30]), c_0_31])])).
fof(c_0_61, plain, ![X9, X10]:(((r1_tarski(X9,X10)|X9!=X10)&(r1_tarski(X10,X9)|X9!=X10))&(~r1_tarski(X9,X10)|~r1_tarski(X10,X9)|X9=X10)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_62, negated_conjecture, (r2_hidden(k4_tarski(esk5_2(esk3_0,X1),X2),esk4_0)|r1_tarski(esk3_0,X1)|X2!=esk6_2(esk3_0,X1)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_63, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_64, negated_conjecture, (esk3_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_65, negated_conjecture, (r2_hidden(k4_tarski(esk5_2(esk3_0,X1),esk6_2(esk3_0,X1)),esk4_0)|r1_tarski(esk3_0,X1)), inference(er,[status(thm)],[c_0_62])).
cnf(c_0_66, negated_conjecture, (~r1_tarski(esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_56]), c_0_64])).
cnf(c_0_67, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_65]), c_0_31])]), c_0_66]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 68
# Proof object clause steps            : 53
# Proof object formula steps           : 15
# Proof object conjectures             : 46
# Proof object clause conjectures      : 43
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 32
# Proof object simplifying inferences  : 44
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 29
# Processed clauses                    : 3561
# ...of these trivial                  : 1
# ...subsumed                          : 991
# ...remaining for further processing  : 2569
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 9735
# ...of the previous two non-trivial   : 9678
# Contextual simplify-reflections      : 30
# Paramodulations                      : 8950
# Factorizations                       : 6
# Equation resolutions                 : 779
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2540
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 2526
# Current number of unprocessed clauses: 6167
# ...number of literals in the above   : 28330
# Current number of archived formulas  : 0
# Current number of archived clauses   : 27
# Clause-clause subsumption calls (NU) : 679474
# Rec. Clause-clause subsumption calls : 141943
# Non-unit clause-clause subsumptions  : 1021
# Unit Clause-clause subsumption calls : 2490
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 300640

# -------------------------------------------------
# User time                : 0.341 s
# System time              : 0.010 s
# Total time               : 0.351 s
# Maximum resident set size: 28748 pages
