# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t26_compts_1, conjecture, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>((v1_compts_1(X1)&(v8_pre_topc(X2)&(k1_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X1)&(k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X2)&(v2_funct_1(X3)&v5_pre_topc(X3,X1,X2))))))=>v3_tops_2(X3,X1,X2))))), file('data/MPTP/proofs/00db599e017543a89b7d0871754527f3.E_input', t26_compts_1)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/00db599e017543a89b7d0871754527f3.E_input', redefinition_k2_relset_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/00db599e017543a89b7d0871754527f3.E_input', redefinition_k1_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/00db599e017543a89b7d0871754527f3.E_input', cc1_relset_1)).
fof(t3_funct_2, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_funct_1(X1)&(v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))))), file('data/MPTP/proofs/00db599e017543a89b7d0871754527f3.E_input', t3_funct_2)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/00db599e017543a89b7d0871754527f3.E_input', d3_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/00db599e017543a89b7d0871754527f3.E_input', dt_l1_pre_topc)).
fof(t72_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:((~(v2_struct_0(X2))&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v3_tops_2(X3,X1,X2)<=>(k1_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X1)&(k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X2)&(v2_funct_1(X3)&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(v4_pre_topc(X4,X1)<=>v4_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X2)))))))))), file('data/MPTP/proofs/00db599e017543a89b7d0871754527f3.E_input', t72_tops_2)).
fof(t164_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((r1_tarski(X1,k1_relat_1(X2))&v2_funct_1(X2))=>k10_relat_1(X2,k9_relat_1(X2,X1))=X1)), file('data/MPTP/proofs/00db599e017543a89b7d0871754527f3.E_input', t164_funct_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/00db599e017543a89b7d0871754527f3.E_input', t3_subset)).
fof(redefinition_k7_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k7_relset_1(X1,X2,X3,X4)=k9_relat_1(X3,X4)), file('data/MPTP/proofs/00db599e017543a89b7d0871754527f3.E_input', redefinition_k7_relset_1)).
fof(redefinition_k8_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k8_relset_1(X1,X2,X3,X4)=k10_relat_1(X3,X4)), file('data/MPTP/proofs/00db599e017543a89b7d0871754527f3.E_input', redefinition_k8_relset_1)).
fof(d12_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v5_pre_topc(X3,X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(v4_pre_topc(X4,X2)=>v4_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X1))))))), file('data/MPTP/proofs/00db599e017543a89b7d0871754527f3.E_input', d12_pre_topc)).
fof(dt_k7_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k7_relset_1(X1,X2,X3,X4),k1_zfmisc_1(X2))), file('data/MPTP/proofs/00db599e017543a89b7d0871754527f3.E_input', dt_k7_relset_1)).
fof(t25_compts_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>((v1_compts_1(X1)&(v8_pre_topc(X2)&(k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X2)&v5_pre_topc(X3,X1,X2))))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(v4_pre_topc(X4,X1)=>v4_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X2))))))), file('data/MPTP/proofs/00db599e017543a89b7d0871754527f3.E_input', t25_compts_1)).
fof(c_0_15, negated_conjecture, ~(![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>((v1_compts_1(X1)&(v8_pre_topc(X2)&(k1_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X1)&(k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X2)&(v2_funct_1(X3)&v5_pre_topc(X3,X1,X2))))))=>v3_tops_2(X3,X1,X2)))))), inference(assume_negation,[status(cth)],[t26_compts_1])).
fof(c_0_16, plain, ![X31, X32, X33]:(~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X32)))|k2_relset_1(X31,X32,X33)=k2_relat_1(X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
fof(c_0_17, negated_conjecture, ((v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0))&((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))))&((v1_compts_1(esk1_0)&(v8_pre_topc(esk2_0)&(k1_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk1_0)&(k2_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)&(v2_funct_1(esk3_0)&v5_pre_topc(esk3_0,esk1_0,esk2_0))))))&~v3_tops_2(esk3_0,esk1_0,esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_15])])])])).
fof(c_0_18, plain, ![X28, X29, X30]:(~m1_subset_1(X30,k1_zfmisc_1(k2_zfmisc_1(X28,X29)))|k1_relset_1(X28,X29,X30)=k1_relat_1(X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
fof(c_0_19, plain, ![X20, X21, X22]:(~m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(X20,X21)))|v1_relat_1(X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_20, plain, ![X27]:((v1_funct_1(X27)|(~v1_relat_1(X27)|~v1_funct_1(X27)))&((v1_funct_2(X27,k1_relat_1(X27),k2_relat_1(X27))|(~v1_relat_1(X27)|~v1_funct_1(X27)))&(m1_subset_1(X27,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X27),k2_relat_1(X27))))|(~v1_relat_1(X27)|~v1_funct_1(X27))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_funct_2])])])).
cnf(c_0_21, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, negated_conjecture, (k2_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, negated_conjecture, (k2_relat_1(esk3_0)=k2_struct_0(esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])).
cnf(c_0_29, negated_conjecture, (k1_relat_1(esk3_0)=k2_struct_0(esk1_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_22]), c_0_25])).
cnf(c_0_30, negated_conjecture, (v1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_26, c_0_22])).
cnf(c_0_31, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_32, plain, ![X34]:(~l1_struct_0(X34)|k2_struct_0(X34)=u1_struct_0(X34)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
fof(c_0_33, plain, ![X40]:(~l1_pre_topc(X40)|l1_struct_0(X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_34, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k2_struct_0(esk1_0),k2_struct_0(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30]), c_0_31])])).
cnf(c_0_35, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_36, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_37, negated_conjecture, (k1_relset_1(k2_struct_0(esk1_0),k2_struct_0(esk2_0),esk3_0)=k2_struct_0(esk1_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_34]), c_0_29])).
cnf(c_0_38, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_39, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_40, negated_conjecture, (k1_relset_1(k2_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])])).
cnf(c_0_41, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_42, negated_conjecture, (k2_relset_1(k2_struct_0(esk1_0),k2_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_34]), c_0_28])).
cnf(c_0_43, negated_conjecture, (k2_struct_0(esk1_0)=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_38]), c_0_25]), c_0_41])])).
fof(c_0_44, plain, ![X35, X36, X37, X38]:(((k1_relset_1(u1_struct_0(X35),u1_struct_0(X36),X37)=k2_struct_0(X35)|~v3_tops_2(X37,X35,X36)|(~v1_funct_1(X37)|(~v1_funct_2(X37,u1_struct_0(X35),u1_struct_0(X36))|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X35),u1_struct_0(X36))))))|(v2_struct_0(X36)|~l1_pre_topc(X36))|~l1_pre_topc(X35))&((k2_relset_1(u1_struct_0(X35),u1_struct_0(X36),X37)=k2_struct_0(X36)|~v3_tops_2(X37,X35,X36)|(~v1_funct_1(X37)|(~v1_funct_2(X37,u1_struct_0(X35),u1_struct_0(X36))|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X35),u1_struct_0(X36))))))|(v2_struct_0(X36)|~l1_pre_topc(X36))|~l1_pre_topc(X35))&((v2_funct_1(X37)|~v3_tops_2(X37,X35,X36)|(~v1_funct_1(X37)|(~v1_funct_2(X37,u1_struct_0(X35),u1_struct_0(X36))|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X35),u1_struct_0(X36))))))|(v2_struct_0(X36)|~l1_pre_topc(X36))|~l1_pre_topc(X35))&((~v4_pre_topc(X38,X35)|v4_pre_topc(k7_relset_1(u1_struct_0(X35),u1_struct_0(X36),X37,X38),X36)|~m1_subset_1(X38,k1_zfmisc_1(u1_struct_0(X35)))|~v3_tops_2(X37,X35,X36)|(~v1_funct_1(X37)|(~v1_funct_2(X37,u1_struct_0(X35),u1_struct_0(X36))|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X35),u1_struct_0(X36))))))|(v2_struct_0(X36)|~l1_pre_topc(X36))|~l1_pre_topc(X35))&(~v4_pre_topc(k7_relset_1(u1_struct_0(X35),u1_struct_0(X36),X37,X38),X36)|v4_pre_topc(X38,X35)|~m1_subset_1(X38,k1_zfmisc_1(u1_struct_0(X35)))|~v3_tops_2(X37,X35,X36)|(~v1_funct_1(X37)|(~v1_funct_2(X37,u1_struct_0(X35),u1_struct_0(X36))|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X35),u1_struct_0(X36))))))|(v2_struct_0(X36)|~l1_pre_topc(X36))|~l1_pre_topc(X35))))))&((m1_subset_1(esk4_3(X35,X36,X37),k1_zfmisc_1(u1_struct_0(X35)))|~v2_funct_1(X37)|k2_relset_1(u1_struct_0(X35),u1_struct_0(X36),X37)!=k2_struct_0(X36)|k1_relset_1(u1_struct_0(X35),u1_struct_0(X36),X37)!=k2_struct_0(X35)|v3_tops_2(X37,X35,X36)|(~v1_funct_1(X37)|(~v1_funct_2(X37,u1_struct_0(X35),u1_struct_0(X36))|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X35),u1_struct_0(X36))))))|(v2_struct_0(X36)|~l1_pre_topc(X36))|~l1_pre_topc(X35))&((~v4_pre_topc(esk4_3(X35,X36,X37),X35)|~v4_pre_topc(k7_relset_1(u1_struct_0(X35),u1_struct_0(X36),X37,esk4_3(X35,X36,X37)),X36)|~v2_funct_1(X37)|k2_relset_1(u1_struct_0(X35),u1_struct_0(X36),X37)!=k2_struct_0(X36)|k1_relset_1(u1_struct_0(X35),u1_struct_0(X36),X37)!=k2_struct_0(X35)|v3_tops_2(X37,X35,X36)|(~v1_funct_1(X37)|(~v1_funct_2(X37,u1_struct_0(X35),u1_struct_0(X36))|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X35),u1_struct_0(X36))))))|(v2_struct_0(X36)|~l1_pre_topc(X36))|~l1_pre_topc(X35))&(v4_pre_topc(esk4_3(X35,X36,X37),X35)|v4_pre_topc(k7_relset_1(u1_struct_0(X35),u1_struct_0(X36),X37,esk4_3(X35,X36,X37)),X36)|~v2_funct_1(X37)|k2_relset_1(u1_struct_0(X35),u1_struct_0(X36),X37)!=k2_struct_0(X36)|k1_relset_1(u1_struct_0(X35),u1_struct_0(X36),X37)!=k2_struct_0(X35)|v3_tops_2(X37,X35,X36)|(~v1_funct_1(X37)|(~v1_funct_2(X37,u1_struct_0(X35),u1_struct_0(X36))|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X35),u1_struct_0(X36))))))|(v2_struct_0(X36)|~l1_pre_topc(X36))|~l1_pre_topc(X35))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t72_tops_2])])])])])])).
cnf(c_0_45, negated_conjecture, (k2_relset_1(u1_struct_0(esk1_0),k2_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)), inference(rw,[status(thm)],[c_0_42, c_0_43])).
fof(c_0_46, plain, ![X8, X9]:(~v1_relat_1(X9)|~v1_funct_1(X9)|(~r1_tarski(X8,k1_relat_1(X9))|~v2_funct_1(X9)|k10_relat_1(X9,k9_relat_1(X9,X8))=X8)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t164_funct_1])])).
fof(c_0_47, plain, ![X10, X11]:((~m1_subset_1(X10,k1_zfmisc_1(X11))|r1_tarski(X10,X11))&(~r1_tarski(X10,X11)|m1_subset_1(X10,k1_zfmisc_1(X11)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_48, plain, (m1_subset_1(esk4_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|v3_tops_2(X3,X1,X2)|v2_struct_0(X2)|~v2_funct_1(X3)|k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)!=k2_struct_0(X2)|k1_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)!=k2_struct_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_49, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)=u1_struct_0(esk1_0)), inference(rw,[status(thm)],[c_0_25, c_0_43])).
cnf(c_0_50, negated_conjecture, (k2_struct_0(esk2_0)=u1_struct_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_38]), c_0_23]), c_0_39])])).
cnf(c_0_51, negated_conjecture, (v2_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_52, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_53, negated_conjecture, (~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_54, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_55, plain, (k10_relat_1(X1,k9_relat_1(X1,X2))=X2|~v1_relat_1(X1)|~v1_funct_1(X1)|~r1_tarski(X2,k1_relat_1(X1))|~v2_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_56, negated_conjecture, (k1_relat_1(esk3_0)=u1_struct_0(esk1_0)), inference(rw,[status(thm)],[c_0_29, c_0_43])).
cnf(c_0_57, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_58, negated_conjecture, (m1_subset_1(esk4_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_43]), c_0_23]), c_0_50]), c_0_50]), c_0_51]), c_0_22]), c_0_52]), c_0_31]), c_0_39]), c_0_41])]), c_0_53]), c_0_54])).
fof(c_0_59, plain, ![X46, X47, X48, X49]:(~m1_subset_1(X48,k1_zfmisc_1(k2_zfmisc_1(X46,X47)))|k7_relset_1(X46,X47,X48,X49)=k9_relat_1(X48,X49)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_relset_1])])).
fof(c_0_60, plain, ![X23, X24, X25, X26]:(~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(X23,X24)))|k8_relset_1(X23,X24,X25,X26)=k10_relat_1(X25,X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_relset_1])])).
fof(c_0_61, plain, ![X41, X42, X43, X44]:((~v5_pre_topc(X43,X41,X42)|(~m1_subset_1(X44,k1_zfmisc_1(u1_struct_0(X42)))|(~v4_pre_topc(X44,X42)|v4_pre_topc(k8_relset_1(u1_struct_0(X41),u1_struct_0(X42),X43,X44),X41)))|(~v1_funct_1(X43)|(~v1_funct_2(X43,u1_struct_0(X41),u1_struct_0(X42))|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X41),u1_struct_0(X42))))))|~l1_pre_topc(X42)|~l1_pre_topc(X41))&((m1_subset_1(esk5_3(X41,X42,X43),k1_zfmisc_1(u1_struct_0(X42)))|v5_pre_topc(X43,X41,X42)|(~v1_funct_1(X43)|(~v1_funct_2(X43,u1_struct_0(X41),u1_struct_0(X42))|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X41),u1_struct_0(X42))))))|~l1_pre_topc(X42)|~l1_pre_topc(X41))&((v4_pre_topc(esk5_3(X41,X42,X43),X42)|v5_pre_topc(X43,X41,X42)|(~v1_funct_1(X43)|(~v1_funct_2(X43,u1_struct_0(X41),u1_struct_0(X42))|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X41),u1_struct_0(X42))))))|~l1_pre_topc(X42)|~l1_pre_topc(X41))&(~v4_pre_topc(k8_relset_1(u1_struct_0(X41),u1_struct_0(X42),X43,esk5_3(X41,X42,X43)),X41)|v5_pre_topc(X43,X41,X42)|(~v1_funct_1(X43)|(~v1_funct_2(X43,u1_struct_0(X41),u1_struct_0(X42))|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X41),u1_struct_0(X42))))))|~l1_pre_topc(X42)|~l1_pre_topc(X41))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_pre_topc])])])])])).
cnf(c_0_62, negated_conjecture, (k10_relat_1(esk3_0,k9_relat_1(esk3_0,X1))=X1|~r1_tarski(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_30]), c_0_51]), c_0_31])])).
cnf(c_0_63, negated_conjecture, (r1_tarski(esk4_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_64, plain, (k7_relset_1(X2,X3,X1,X4)=k9_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_65, plain, (k8_relset_1(X2,X3,X1,X4)=k10_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_60])).
fof(c_0_66, plain, ![X16, X17, X18, X19]:(~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(X16,X17)))|m1_subset_1(k7_relset_1(X16,X17,X18,X19),k1_zfmisc_1(X17))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_relset_1])])).
cnf(c_0_67, plain, (v4_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(X3),X1,X4),X2)|~v5_pre_topc(X1,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~v4_pre_topc(X4,X3)|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(X3)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_68, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_69, plain, (v4_pre_topc(esk4_3(X1,X2,X3),X1)|v4_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk4_3(X1,X2,X3)),X2)|v3_tops_2(X3,X1,X2)|v2_struct_0(X2)|~v2_funct_1(X3)|k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)!=k2_struct_0(X2)|k1_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)!=k2_struct_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_70, negated_conjecture, (k10_relat_1(esk3_0,k9_relat_1(esk3_0,esk4_3(esk1_0,esk2_0,esk3_0)))=esk4_3(esk1_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_71, negated_conjecture, (k9_relat_1(esk3_0,X1)=k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1)), inference(spm,[status(thm)],[c_0_64, c_0_22])).
cnf(c_0_72, negated_conjecture, (k10_relat_1(esk3_0,X1)=k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1)), inference(spm,[status(thm)],[c_0_65, c_0_22])).
cnf(c_0_73, plain, (m1_subset_1(k7_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_66])).
fof(c_0_74, plain, ![X12, X13, X14, X15]:(~v2_pre_topc(X12)|~l1_pre_topc(X12)|(v2_struct_0(X13)|(~v2_pre_topc(X13)|~l1_pre_topc(X13))|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(X12),u1_struct_0(X13))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X12),u1_struct_0(X13)))))|(~v1_compts_1(X12)|(~v8_pre_topc(X13)|(k2_relset_1(u1_struct_0(X12),u1_struct_0(X13),X14)!=k2_struct_0(X13)|~v5_pre_topc(X14,X12,X13)))|(~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X12)))|(~v4_pre_topc(X15,X12)|v4_pre_topc(k7_relset_1(u1_struct_0(X12),u1_struct_0(X13),X14,X15),X13))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t25_compts_1])])])])).
cnf(c_0_75, plain, (v3_tops_2(X3,X1,X2)|v2_struct_0(X2)|~v4_pre_topc(esk4_3(X1,X2,X3),X1)|~v4_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk4_3(X1,X2,X3)),X2)|~v2_funct_1(X3)|k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)!=k2_struct_0(X2)|k1_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)!=k2_struct_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_76, negated_conjecture, (v4_pre_topc(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1),esk1_0)|~v4_pre_topc(X1,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_22]), c_0_52]), c_0_31]), c_0_39]), c_0_41])])).
cnf(c_0_77, negated_conjecture, (v4_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_3(esk1_0,esk2_0,esk3_0)),esk2_0)|v4_pre_topc(esk4_3(esk1_0,esk2_0,esk3_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_49]), c_0_43]), c_0_23]), c_0_50]), c_0_50]), c_0_51]), c_0_22]), c_0_52]), c_0_31]), c_0_39]), c_0_41])]), c_0_53]), c_0_54])).
cnf(c_0_78, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_3(esk1_0,esk2_0,esk3_0)))=esk4_3(esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_71]), c_0_72])).
cnf(c_0_79, negated_conjecture, (m1_subset_1(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_73, c_0_22])).
cnf(c_0_80, plain, (v2_struct_0(X2)|v4_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_compts_1(X1)|~v8_pre_topc(X2)|k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v5_pre_topc(X3,X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v4_pre_topc(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_81, negated_conjecture, (v8_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_82, negated_conjecture, (v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_83, negated_conjecture, (v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_84, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_85, negated_conjecture, (~v4_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~v4_pre_topc(esk4_3(esk1_0,esk2_0,esk3_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_49]), c_0_43]), c_0_23]), c_0_50]), c_0_50]), c_0_51]), c_0_22]), c_0_52]), c_0_31]), c_0_39]), c_0_41])]), c_0_53]), c_0_54])).
cnf(c_0_86, negated_conjecture, (v4_pre_topc(esk4_3(esk1_0,esk2_0,esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78]), c_0_79])])).
cnf(c_0_87, negated_conjecture, (v4_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1),esk2_0)|~v4_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_68]), c_0_23]), c_0_50]), c_0_50]), c_0_81]), c_0_82]), c_0_22]), c_0_52]), c_0_31]), c_0_39]), c_0_41]), c_0_83]), c_0_84])]), c_0_54])).
cnf(c_0_88, negated_conjecture, (~v4_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_3(esk1_0,esk2_0,esk3_0)),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_86])])).
cnf(c_0_89, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_86]), c_0_58])]), c_0_88]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 90
# Proof object clause steps            : 59
# Proof object formula steps           : 31
# Proof object conjectures             : 45
# Proof object clause conjectures      : 42
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 23
# Proof object simplifying inferences  : 92
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 42
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 41
# Processed clauses                    : 125
# ...of these trivial                  : 3
# ...subsumed                          : 0
# ...remaining for further processing  : 121
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 16
# Generated clauses                    : 67
# ...of the previous two non-trivial   : 58
# Contextual simplify-reflections      : 0
# Paramodulations                      : 67
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 64
#    Positive orientable unit clauses  : 27
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 34
# Current number of unprocessed clauses: 3
# ...number of literals in the above   : 9
# Current number of archived formulas  : 0
# Current number of archived clauses   : 57
# Clause-clause subsumption calls (NU) : 2060
# Rec. Clause-clause subsumption calls : 74
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 5
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 10
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6349

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.000 s
# Total time               : 0.013 s
# Maximum resident set size: 28796 pages
