# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C47_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t40_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>((r1_tsep_1(k1_tsep_1(X1,X2,X3),X4)=>(r1_tsep_1(X2,X4)&r1_tsep_1(X3,X4)))&(((r1_tsep_1(X2,X4)&r1_tsep_1(X3,X4))=>r1_tsep_1(k1_tsep_1(X1,X2,X3),X4))&((r1_tsep_1(X4,k1_tsep_1(X1,X2,X3))=>(r1_tsep_1(X4,X2)&r1_tsep_1(X4,X3)))&((r1_tsep_1(X4,X2)&r1_tsep_1(X4,X3))=>r1_tsep_1(X4,k1_tsep_1(X1,X2,X3)))))))))), file('data/MPTP/proofs/00ef166038074bea8e78555fdc86829c.E_input', t40_tmap_1)).
fof(t39_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(~((~(r1_tsep_1(k1_tsep_1(X1,X2,X3),X4))&(r1_tsep_1(X2,X4)&r1_tsep_1(X3,X4))))&(~((~((r1_tsep_1(X2,X4)&r1_tsep_1(X3,X4)))&r1_tsep_1(k1_tsep_1(X1,X2,X3),X4)))&(~((~(r1_tsep_1(X4,k1_tsep_1(X1,X2,X3)))&(r1_tsep_1(X4,X2)&r1_tsep_1(X4,X3))))&~((~((r1_tsep_1(X4,X2)&r1_tsep_1(X4,X3)))&r1_tsep_1(X4,k1_tsep_1(X1,X2,X3))))))))))), file('data/MPTP/proofs/00ef166038074bea8e78555fdc86829c.E_input', t39_tmap_1)).
fof(c_0_2, plain, ![X1, X2, X3, X4]:(epred1_4(X4,X2,X3,X1)<=>(~((~(r1_tsep_1(k1_tsep_1(X1,X2,X3),X4))&(r1_tsep_1(X2,X4)&r1_tsep_1(X3,X4))))&(~((~((r1_tsep_1(X2,X4)&r1_tsep_1(X3,X4)))&r1_tsep_1(k1_tsep_1(X1,X2,X3),X4)))&(~((~(r1_tsep_1(X4,k1_tsep_1(X1,X2,X3)))&(r1_tsep_1(X4,X2)&r1_tsep_1(X4,X3))))&~((~((r1_tsep_1(X4,X2)&r1_tsep_1(X4,X3)))&r1_tsep_1(X4,k1_tsep_1(X1,X2,X3)))))))), introduced(definition)).
fof(c_0_3, plain, ![X1, X2, X3, X4]:(epred1_4(X4,X2,X3,X1)=>(~((~(r1_tsep_1(k1_tsep_1(X1,X2,X3),X4))&(r1_tsep_1(X2,X4)&r1_tsep_1(X3,X4))))&(~((~((r1_tsep_1(X2,X4)&r1_tsep_1(X3,X4)))&r1_tsep_1(k1_tsep_1(X1,X2,X3),X4)))&(~((~(r1_tsep_1(X4,k1_tsep_1(X1,X2,X3)))&(r1_tsep_1(X4,X2)&r1_tsep_1(X4,X3))))&~((~((r1_tsep_1(X4,X2)&r1_tsep_1(X4,X3)))&r1_tsep_1(X4,k1_tsep_1(X1,X2,X3)))))))), inference(split_equiv,[status(thm)],[c_0_2])).
fof(c_0_4, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>((r1_tsep_1(k1_tsep_1(X1,X2,X3),X4)=>(r1_tsep_1(X2,X4)&r1_tsep_1(X3,X4)))&(((r1_tsep_1(X2,X4)&r1_tsep_1(X3,X4))=>r1_tsep_1(k1_tsep_1(X1,X2,X3),X4))&((r1_tsep_1(X4,k1_tsep_1(X1,X2,X3))=>(r1_tsep_1(X4,X2)&r1_tsep_1(X4,X3)))&((r1_tsep_1(X4,X2)&r1_tsep_1(X4,X3))=>r1_tsep_1(X4,k1_tsep_1(X1,X2,X3))))))))))), inference(assume_negation,[status(cth)],[t40_tmap_1])).
fof(c_0_5, plain, ![X13, X14, X15, X16]:((r1_tsep_1(k1_tsep_1(X13,X14,X15),X16)|(~r1_tsep_1(X14,X16)|~r1_tsep_1(X15,X16))|~epred1_4(X16,X14,X15,X13))&(((r1_tsep_1(X14,X16)|~r1_tsep_1(k1_tsep_1(X13,X14,X15),X16)|~epred1_4(X16,X14,X15,X13))&(r1_tsep_1(X15,X16)|~r1_tsep_1(k1_tsep_1(X13,X14,X15),X16)|~epred1_4(X16,X14,X15,X13)))&((r1_tsep_1(X16,k1_tsep_1(X13,X14,X15))|(~r1_tsep_1(X16,X14)|~r1_tsep_1(X16,X15))|~epred1_4(X16,X14,X15,X13))&((r1_tsep_1(X16,X14)|~r1_tsep_1(X16,k1_tsep_1(X13,X14,X15))|~epred1_4(X16,X14,X15,X13))&(r1_tsep_1(X16,X15)|~r1_tsep_1(X16,k1_tsep_1(X13,X14,X15))|~epred1_4(X16,X14,X15,X13)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_3])])])])).
fof(c_0_6, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk1_0))&((~v2_struct_0(esk4_0)&m1_pre_topc(esk4_0,esk1_0))&(((((((r1_tsep_1(esk4_0,esk2_0)|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk2_0,esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0))&(r1_tsep_1(esk4_0,esk3_0)|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk2_0,esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)))&(~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk2_0,esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)))&(((r1_tsep_1(esk4_0,esk2_0)|(~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0))|r1_tsep_1(esk2_0,esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0))&(r1_tsep_1(esk4_0,esk3_0)|(~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0))|r1_tsep_1(esk2_0,esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)))&(~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|(~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0))|r1_tsep_1(esk2_0,esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0))))&((((r1_tsep_1(esk4_0,esk2_0)|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk3_0,esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0))&(r1_tsep_1(esk4_0,esk3_0)|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk3_0,esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)))&(~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk3_0,esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)))&(((r1_tsep_1(esk4_0,esk2_0)|(~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0))|r1_tsep_1(esk3_0,esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0))&(r1_tsep_1(esk4_0,esk3_0)|(~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0))|r1_tsep_1(esk3_0,esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)))&(~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|(~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0))|r1_tsep_1(esk3_0,esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)))))&((((r1_tsep_1(esk4_0,esk2_0)|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0))&(r1_tsep_1(esk4_0,esk3_0)|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)))&(~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)))&(((r1_tsep_1(esk4_0,esk2_0)|(~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0))&(r1_tsep_1(esk4_0,esk3_0)|(~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)))&(~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|(~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)))))&((((((r1_tsep_1(esk4_0,esk2_0)|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk2_0,esk4_0)|(~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0)))&(r1_tsep_1(esk4_0,esk3_0)|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk2_0,esk4_0)|(~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0))))&(~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk2_0,esk4_0)|(~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0))))&(((r1_tsep_1(esk4_0,esk2_0)|(~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0))|r1_tsep_1(esk2_0,esk4_0)|(~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0)))&(r1_tsep_1(esk4_0,esk3_0)|(~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0))|r1_tsep_1(esk2_0,esk4_0)|(~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0))))&(~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|(~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0))|r1_tsep_1(esk2_0,esk4_0)|(~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0)))))&((((r1_tsep_1(esk4_0,esk2_0)|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk3_0,esk4_0)|(~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0)))&(r1_tsep_1(esk4_0,esk3_0)|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk3_0,esk4_0)|(~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0))))&(~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk3_0,esk4_0)|(~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0))))&(((r1_tsep_1(esk4_0,esk2_0)|(~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0))|r1_tsep_1(esk3_0,esk4_0)|(~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0)))&(r1_tsep_1(esk4_0,esk3_0)|(~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0))|r1_tsep_1(esk3_0,esk4_0)|(~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0))))&(~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|(~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0))|r1_tsep_1(esk3_0,esk4_0)|(~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0))))))&((((r1_tsep_1(esk4_0,esk2_0)|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|(~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0)))&(r1_tsep_1(esk4_0,esk3_0)|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|(~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0))))&(~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|(~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0))))&(((r1_tsep_1(esk4_0,esk2_0)|(~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|(~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0)))&(r1_tsep_1(esk4_0,esk3_0)|(~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|(~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0))))&(~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|(~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|(~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0))))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_4])])])])])).
cnf(c_0_7, plain, (r1_tsep_1(X1,X2)|~r1_tsep_1(k1_tsep_1(X3,X1,X4),X2)|~epred1_4(X2,X1,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8, negated_conjecture, (r1_tsep_1(esk2_0,esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, plain, (r1_tsep_1(X1,X2)|~r1_tsep_1(X1,k1_tsep_1(X3,X4,X2))|~epred1_4(X1,X4,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10, plain, (r1_tsep_1(X1,X2)|~r1_tsep_1(X1,k1_tsep_1(X3,X2,X4))|~epred1_4(X1,X2,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11, negated_conjecture, (~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12, plain, (r1_tsep_1(k1_tsep_1(X1,X2,X3),X4)|~r1_tsep_1(X2,X4)|~r1_tsep_1(X3,X4)|~epred1_4(X4,X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_13, negated_conjecture, (r1_tsep_1(esk2_0,esk4_0)|~epred1_4(esk4_0,esk2_0,esk3_0,esk1_0)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_9]), c_0_10])).
cnf(c_0_14, plain, (r1_tsep_1(X1,X2)|~r1_tsep_1(k1_tsep_1(X3,X4,X1),X2)|~epred1_4(X2,X4,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_15, negated_conjecture, (r1_tsep_1(esk3_0,esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_16, negated_conjecture, (~epred1_4(esk4_0,esk2_0,esk3_0,esk1_0)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|~r1_tsep_1(esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_9]), c_0_10]), c_0_13])).
fof(c_0_17, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>epred1_4(X4,X2,X3,X1))))), inference(apply_def,[status(thm)],[t39_tmap_1, c_0_2])).
cnf(c_0_18, plain, (~epred1_4(esk4_0,esk2_0,esk3_0,esk1_0)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_9]), c_0_10]), c_0_16])).
cnf(c_0_19, plain, (r1_tsep_1(X1,k1_tsep_1(X2,X3,X4))|~r1_tsep_1(X1,X3)|~r1_tsep_1(X1,X4)|~epred1_4(X1,X3,X4,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_20, plain, ![X9, X10, X11, X12]:(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))|(v2_struct_0(X10)|~m1_pre_topc(X10,X9)|(v2_struct_0(X11)|~m1_pre_topc(X11,X9)|(v2_struct_0(X12)|~m1_pre_topc(X12,X9)|epred1_4(X12,X10,X11,X9))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])).
cnf(c_0_21, negated_conjecture, (r1_tsep_1(esk4_0,esk3_0)|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk2_0,esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_22, negated_conjecture, (r1_tsep_1(esk4_0,esk2_0)|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk2_0,esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_23, plain, (~epred1_4(esk4_0,esk2_0,esk3_0,esk1_0)|~r1_tsep_1(esk4_0,esk3_0)|~r1_tsep_1(esk4_0,esk2_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_24, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|epred1_4(X4,X2,X3,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25, negated_conjecture, (m1_pre_topc(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_26, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_27, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_28, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_29, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_30, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_31, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_32, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_33, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_34, plain, (r1_tsep_1(esk4_0,esk3_0)|r1_tsep_1(esk2_0,esk4_0)|~epred1_4(esk4_0,esk2_0,esk3_0,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_21]), c_0_9])).
cnf(c_0_35, negated_conjecture, (r1_tsep_1(esk4_0,esk2_0)|r1_tsep_1(esk2_0,esk4_0)|~epred1_4(esk4_0,esk2_0,esk3_0,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_22]), c_0_10])).
cnf(c_0_36, plain, (~r1_tsep_1(esk4_0,esk3_0)|~r1_tsep_1(esk4_0,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29])]), c_0_30]), c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_37, plain, (r1_tsep_1(esk2_0,esk4_0)|r1_tsep_1(esk4_0,esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29])]), c_0_30]), c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_38, negated_conjecture, (r1_tsep_1(esk2_0,esk4_0)|r1_tsep_1(esk4_0,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29])]), c_0_30]), c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_39, negated_conjecture, (r1_tsep_1(esk4_0,esk3_0)|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_40, plain, (r1_tsep_1(esk2_0,esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_41, negated_conjecture, (r1_tsep_1(esk4_0,esk2_0)|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk3_0,esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_42, negated_conjecture, (r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk4_0,esk3_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|~r1_tsep_1(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_40])])).
cnf(c_0_43, negated_conjecture, (r1_tsep_1(esk4_0,esk2_0)|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_44, plain, (r1_tsep_1(esk4_0,esk2_0)|r1_tsep_1(esk3_0,esk4_0)|~epred1_4(esk4_0,esk2_0,esk3_0,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_41]), c_0_10])).
cnf(c_0_45, plain, (r1_tsep_1(esk4_0,esk3_0)|~epred1_4(esk4_0,esk2_0,esk3_0,esk1_0)|~r1_tsep_1(esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_12]), c_0_40])]), c_0_9])).
cnf(c_0_46, negated_conjecture, (r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|~r1_tsep_1(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_40])])).
cnf(c_0_47, plain, (r1_tsep_1(esk3_0,esk4_0)|r1_tsep_1(esk4_0,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29])]), c_0_30]), c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_48, plain, (r1_tsep_1(esk4_0,esk3_0)|~r1_tsep_1(esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29])]), c_0_30]), c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_49, plain, (r1_tsep_1(esk4_0,esk2_0)|~epred1_4(esk4_0,esk2_0,esk3_0,esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_12]), c_0_40])]), c_0_47]), c_0_10])).
cnf(c_0_50, plain, (~r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_36, c_0_48])).
cnf(c_0_51, plain, (r1_tsep_1(esk4_0,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29])]), c_0_30]), c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_52, negated_conjecture, (r1_tsep_1(esk4_0,esk3_0)|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk3_0,esk4_0)|r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_53, plain, (~r1_tsep_1(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_51])])).
cnf(c_0_54, plain, (~r1_tsep_1(esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_51])])).
cnf(c_0_55, negated_conjecture, (r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_52, c_0_53]), c_0_54])).
cnf(c_0_56, plain, (~epred1_4(esk4_0,esk2_0,esk3_0,esk1_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_55]), c_0_53]), c_0_18])).
cnf(c_0_57, plain, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29])]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 58
# Proof object clause steps            : 49
# Proof object formula steps           : 9
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 2
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 101
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 2
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 52
# Removed in clause preprocessing      : 27
# Initial clauses in saturation        : 25
# Processed clauses                    : 84
# ...of these trivial                  : 1
# ...subsumed                          : 2
# ...remaining for further processing  : 81
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 19
# Backward-rewritten                   : 8
# Generated clauses                    : 70
# ...of the previous two non-trivial   : 45
# Contextual simplify-reflections      : 16
# Paramodulations                      : 69
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 28
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 14
# Current number of unprocessed clauses: 8
# ...number of literals in the above   : 71
# Current number of archived formulas  : 0
# Current number of archived clauses   : 53
# Clause-clause subsumption calls (NU) : 458
# Rec. Clause-clause subsumption calls : 169
# Non-unit clause-clause subsumptions  : 35
# Unit Clause-clause subsumption calls : 10
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4021

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28720 pages
