# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t37_wellord1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>((v2_wellord1(X3)&(r2_hidden(X1,k3_relat_1(X3))&r2_hidden(X2,k3_relat_1(X3))))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>r1_tarski(k1_wellord1(X3,X1),k1_wellord1(X3,X2))))), file('data/MPTP/proofs/00f6c403823941dcb1f4cee7ce444d0c.E_input', t37_wellord1)).
fof(l4_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v6_relat_2(X1)<=>![X2, X3]:~((r2_hidden(X2,k3_relat_1(X1))&(r2_hidden(X3,k3_relat_1(X1))&(~(X2=X3)&(~(r2_hidden(k4_tarski(X2,X3),X1))&~(r2_hidden(k4_tarski(X3,X2),X1))))))))), file('data/MPTP/proofs/00f6c403823941dcb1f4cee7ce444d0c.E_input', l4_wellord1)).
fof(d4_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v2_wellord1(X1)<=>(v1_relat_2(X1)&(v8_relat_2(X1)&(v4_relat_2(X1)&(v6_relat_2(X1)&v1_wellord1(X1))))))), file('data/MPTP/proofs/00f6c403823941dcb1f4cee7ce444d0c.E_input', d4_wellord1)).
fof(d1_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(X3=k1_wellord1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(~(X4=X2)&r2_hidden(k4_tarski(X4,X2),X1))))), file('data/MPTP/proofs/00f6c403823941dcb1f4cee7ce444d0c.E_input', d1_wellord1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/00f6c403823941dcb1f4cee7ce444d0c.E_input', d3_tarski)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/00f6c403823941dcb1f4cee7ce444d0c.E_input', t1_xboole_1)).
fof(t13_wellord1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k1_wellord1(X2,X1),k3_relat_1(X2))), file('data/MPTP/proofs/00f6c403823941dcb1f4cee7ce444d0c.E_input', t13_wellord1)).
fof(t35_wellord1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>((v2_wellord1(X3)&r2_hidden(X1,k1_wellord1(X3,X2)))=>k1_wellord1(k2_wellord1(X3,k1_wellord1(X3,X2)),X1)=k1_wellord1(X3,X1))), file('data/MPTP/proofs/00f6c403823941dcb1f4cee7ce444d0c.E_input', t35_wellord1)).
fof(dt_k2_wellord1, axiom, ![X1, X2]:(v1_relat_1(X1)=>v1_relat_1(k2_wellord1(X1,X2))), file('data/MPTP/proofs/00f6c403823941dcb1f4cee7ce444d0c.E_input', dt_k2_wellord1)).
fof(t20_wellord1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(k3_relat_1(k2_wellord1(X2,X1)),k3_relat_1(X2))&r1_tarski(k3_relat_1(k2_wellord1(X2,X1)),X1))), file('data/MPTP/proofs/00f6c403823941dcb1f4cee7ce444d0c.E_input', t20_wellord1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/00f6c403823941dcb1f4cee7ce444d0c.E_input', d10_xboole_0)).
fof(l1_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v1_relat_2(X1)<=>![X2]:(r2_hidden(X2,k3_relat_1(X1))=>r2_hidden(k4_tarski(X2,X2),X1)))), file('data/MPTP/proofs/00f6c403823941dcb1f4cee7ce444d0c.E_input', l1_wellord1)).
fof(c_0_12, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>((v2_wellord1(X3)&(r2_hidden(X1,k3_relat_1(X3))&r2_hidden(X2,k3_relat_1(X3))))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>r1_tarski(k1_wellord1(X3,X1),k1_wellord1(X3,X2)))))), inference(assume_negation,[status(cth)],[t37_wellord1])).
fof(c_0_13, plain, ![X31, X32, X33]:((~v6_relat_2(X31)|(~r2_hidden(X32,k3_relat_1(X31))|(~r2_hidden(X33,k3_relat_1(X31))|(X32=X33|(r2_hidden(k4_tarski(X32,X33),X31)|r2_hidden(k4_tarski(X33,X32),X31)))))|~v1_relat_1(X31))&((r2_hidden(esk6_1(X31),k3_relat_1(X31))|v6_relat_2(X31)|~v1_relat_1(X31))&((r2_hidden(esk7_1(X31),k3_relat_1(X31))|v6_relat_2(X31)|~v1_relat_1(X31))&((esk6_1(X31)!=esk7_1(X31)|v6_relat_2(X31)|~v1_relat_1(X31))&((~r2_hidden(k4_tarski(esk6_1(X31),esk7_1(X31)),X31)|v6_relat_2(X31)|~v1_relat_1(X31))&(~r2_hidden(k4_tarski(esk7_1(X31),esk6_1(X31)),X31)|v6_relat_2(X31)|~v1_relat_1(X31))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l4_wellord1])])])])])])])).
fof(c_0_14, negated_conjecture, (v1_relat_1(esk3_0)&((v2_wellord1(esk3_0)&(r2_hidden(esk1_0,k3_relat_1(esk3_0))&r2_hidden(esk2_0,k3_relat_1(esk3_0))))&((~r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|~r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0)))&(r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
cnf(c_0_15, plain, (X2=X3|r2_hidden(k4_tarski(X2,X3),X1)|r2_hidden(k4_tarski(X3,X2),X1)|~v6_relat_2(X1)|~r2_hidden(X2,k3_relat_1(X1))|~r2_hidden(X3,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_16, negated_conjecture, (r2_hidden(esk2_0,k3_relat_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_17, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, plain, ![X23]:(((v1_relat_2(X23)|~v2_wellord1(X23)|~v1_relat_1(X23))&((v8_relat_2(X23)|~v2_wellord1(X23)|~v1_relat_1(X23))&((v4_relat_2(X23)|~v2_wellord1(X23)|~v1_relat_1(X23))&((v6_relat_2(X23)|~v2_wellord1(X23)|~v1_relat_1(X23))&(v1_wellord1(X23)|~v2_wellord1(X23)|~v1_relat_1(X23))))))&(~v1_relat_2(X23)|(~v8_relat_2(X23)|(~v4_relat_2(X23)|(~v6_relat_2(X23)|~v1_wellord1(X23))))|v2_wellord1(X23)|~v1_relat_1(X23))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_wellord1])])])).
fof(c_0_19, plain, ![X8, X9, X10, X11, X12, X13, X14]:((((X11!=X9|~r2_hidden(X11,X10)|X10!=k1_wellord1(X8,X9)|~v1_relat_1(X8))&(r2_hidden(k4_tarski(X11,X9),X8)|~r2_hidden(X11,X10)|X10!=k1_wellord1(X8,X9)|~v1_relat_1(X8)))&(X12=X9|~r2_hidden(k4_tarski(X12,X9),X8)|r2_hidden(X12,X10)|X10!=k1_wellord1(X8,X9)|~v1_relat_1(X8)))&((~r2_hidden(esk4_3(X8,X13,X14),X14)|(esk4_3(X8,X13,X14)=X13|~r2_hidden(k4_tarski(esk4_3(X8,X13,X14),X13),X8))|X14=k1_wellord1(X8,X13)|~v1_relat_1(X8))&((esk4_3(X8,X13,X14)!=X13|r2_hidden(esk4_3(X8,X13,X14),X14)|X14=k1_wellord1(X8,X13)|~v1_relat_1(X8))&(r2_hidden(k4_tarski(esk4_3(X8,X13,X14),X13),X8)|r2_hidden(esk4_3(X8,X13,X14),X14)|X14=k1_wellord1(X8,X13)|~v1_relat_1(X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_wellord1])])])])])])])).
cnf(c_0_20, negated_conjecture, (X1=esk2_0|r2_hidden(k4_tarski(X1,esk2_0),esk3_0)|r2_hidden(k4_tarski(esk2_0,X1),esk3_0)|~v6_relat_2(esk3_0)|~r2_hidden(X1,k3_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])])).
cnf(c_0_21, plain, (v6_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, negated_conjecture, (v2_wellord1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23, plain, (X1=X2|r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_wellord1(X3,X2)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, negated_conjecture, (X1=esk2_0|r2_hidden(k4_tarski(esk2_0,X1),esk3_0)|r2_hidden(k4_tarski(X1,esk2_0),esk3_0)|~r2_hidden(X1,k3_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_17])])).
cnf(c_0_25, negated_conjecture, (r2_hidden(esk1_0,k3_relat_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_26, plain, ![X36, X37, X38, X39, X40]:((~r1_tarski(X36,X37)|(~r2_hidden(X38,X36)|r2_hidden(X38,X37)))&((r2_hidden(esk8_2(X39,X40),X39)|r1_tarski(X39,X40))&(~r2_hidden(esk8_2(X39,X40),X40)|r1_tarski(X39,X40)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_27, plain, (X1=X2|r2_hidden(X1,k1_wellord1(X3,X2))|~r2_hidden(k4_tarski(X1,X2),X3)|~v1_relat_1(X3)), inference(er,[status(thm)],[c_0_23])).
cnf(c_0_28, negated_conjecture, (esk2_0=esk1_0|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r2_hidden(k4_tarski(esk2_0,esk1_0),esk3_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_30, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_31, negated_conjecture, (esk2_0=esk1_0|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_17])])).
fof(c_0_32, plain, ![X42, X43, X44]:(~r1_tarski(X42,X43)|~r1_tarski(X43,X44)|r1_tarski(X42,X44)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_33, plain, ![X21, X22]:(~v1_relat_1(X22)|r1_tarski(k1_wellord1(X22,X21),k3_relat_1(X22))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_wellord1])])).
cnf(c_0_34, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r2_hidden(X1,k1_wellord1(esk3_0,esk2_0))|~r2_hidden(X1,k1_wellord1(esk3_0,esk1_0))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_35, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))|r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_31]), c_0_17])])).
cnf(c_0_36, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_37, plain, (r1_tarski(k1_wellord1(X1,X2),k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_38, plain, ![X18, X19, X20]:(~v1_relat_1(X20)|(~v2_wellord1(X20)|~r2_hidden(X18,k1_wellord1(X20,X19))|k1_wellord1(k2_wellord1(X20,k1_wellord1(X20,X19)),X18)=k1_wellord1(X20,X18))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t35_wellord1])])).
fof(c_0_39, plain, ![X24, X25]:(~v1_relat_1(X24)|v1_relat_1(k2_wellord1(X24,X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_wellord1])])).
cnf(c_0_40, plain, (X1!=X2|~r2_hidden(X1,X3)|X3!=k1_wellord1(X4,X2)|~v1_relat_1(X4)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_41, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk2_0))|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
fof(c_0_42, plain, ![X29, X30]:((r1_tarski(k3_relat_1(k2_wellord1(X30,X29)),k3_relat_1(X30))|~v1_relat_1(X30))&(r1_tarski(k3_relat_1(k2_wellord1(X30,X29)),X29)|~v1_relat_1(X30))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t20_wellord1])])])).
cnf(c_0_43, plain, (r1_tarski(X1,k3_relat_1(X2))|~r1_tarski(X1,k1_wellord1(X2,X3))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_44, plain, (k1_wellord1(k2_wellord1(X1,k1_wellord1(X1,X3)),X2)=k1_wellord1(X1,X2)|~v1_relat_1(X1)|~v2_wellord1(X1)|~r2_hidden(X2,k1_wellord1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_45, plain, (v1_relat_1(k2_wellord1(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_46, plain, (~r2_hidden(X1,k1_wellord1(X2,X1))|~v1_relat_1(X2)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_40])])).
cnf(c_0_47, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk2_0))|r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_41]), c_0_17])])).
cnf(c_0_48, plain, (r1_tarski(k3_relat_1(k2_wellord1(X1,X2)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_49, plain, (r1_tarski(X1,k3_relat_1(k2_wellord1(X2,k1_wellord1(X2,X3))))|~r1_tarski(X1,k1_wellord1(X2,X4))|~r2_hidden(X4,k1_wellord1(X2,X3))|~v2_wellord1(X2)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])).
cnf(c_0_50, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_17])])).
fof(c_0_51, plain, ![X16, X17]:(((r1_tarski(X16,X17)|X16!=X17)&(r1_tarski(X17,X16)|X16!=X17))&(~r1_tarski(X16,X17)|~r1_tarski(X17,X16)|X16=X17)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_52, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k3_relat_1(k2_wellord1(X3,X2)))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_36, c_0_48])).
cnf(c_0_53, negated_conjecture, (esk2_0=esk1_0|r1_tarski(X1,k3_relat_1(k2_wellord1(esk3_0,k1_wellord1(esk3_0,esk2_0))))|~r1_tarski(X1,k1_wellord1(esk3_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_22]), c_0_17])])).
cnf(c_0_54, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_55, negated_conjecture, (~r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|~r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_56, negated_conjecture, (esk2_0=esk1_0|r1_tarski(X1,k1_wellord1(esk3_0,esk2_0))|~r1_tarski(X1,k1_wellord1(esk3_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_17])])).
cnf(c_0_57, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_54])).
cnf(c_0_58, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,X4)|X4!=k1_wellord1(X3,X2)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_59, negated_conjecture, (esk2_0=esk1_0|~r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57])])).
cnf(c_0_60, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,k1_wellord1(X3,X2))|~v1_relat_1(X3)), inference(er,[status(thm)],[c_0_58])).
fof(c_0_61, plain, ![X26, X27]:((~v1_relat_2(X26)|(~r2_hidden(X27,k3_relat_1(X26))|r2_hidden(k4_tarski(X27,X27),X26))|~v1_relat_1(X26))&((r2_hidden(esk5_1(X26),k3_relat_1(X26))|v1_relat_2(X26)|~v1_relat_1(X26))&(~r2_hidden(k4_tarski(esk5_1(X26),esk5_1(X26)),X26)|v1_relat_2(X26)|~v1_relat_1(X26)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_wellord1])])])])])).
cnf(c_0_62, negated_conjecture, (esk2_0=esk1_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_17])]), c_0_50])).
cnf(c_0_63, plain, (r2_hidden(k4_tarski(X2,X2),X1)|~v1_relat_2(X1)|~r2_hidden(X2,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_64, plain, (v1_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_65, negated_conjecture, (~r2_hidden(k4_tarski(esk1_0,esk1_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_62]), c_0_57]), c_0_62])])).
cnf(c_0_66, plain, (r2_hidden(k4_tarski(X1,X1),X2)|~r2_hidden(X1,k3_relat_1(X2))|~v2_wellord1(X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_67, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_25]), c_0_22]), c_0_17])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 68
# Proof object clause steps            : 43
# Proof object formula steps           : 25
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 37
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 39
# Processed clauses                    : 450
# ...of these trivial                  : 0
# ...subsumed                          : 159
# ...remaining for further processing  : 291
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 102
# Generated clauses                    : 1018
# ...of the previous two non-trivial   : 1018
# Contextual simplify-reflections      : 15
# Paramodulations                      : 1013
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 140
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 134
# Current number of unprocessed clauses: 644
# ...number of literals in the above   : 5524
# Current number of archived formulas  : 0
# Current number of archived clauses   : 146
# Clause-clause subsumption calls (NU) : 33386
# Rec. Clause-clause subsumption calls : 1748
# Non-unit clause-clause subsumptions  : 180
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 137555

# -------------------------------------------------
# User time                : 0.073 s
# System time              : 0.003 s
# Total time               : 0.076 s
# Maximum resident set size: 28700 pages
