# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t20_tex_2, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_tex_2(k1_tex_2(X1,X2),X1)<=>v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1))))), file('data/MPTP/proofs/0144bb96470949329693c24e483a6486.E_input', t20_tex_2)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/0144bb96470949329693c24e483a6486.E_input', dt_l1_pre_topc)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/0144bb96470949329693c24e483a6486.E_input', fc2_struct_0)).
fof(dt_k1_tex_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>(~(v2_struct_0(k1_tex_2(X1,X2)))&(v1_pre_topc(k1_tex_2(X1,X2))&m1_pre_topc(k1_tex_2(X1,X2),X1)))), file('data/MPTP/proofs/0144bb96470949329693c24e483a6486.E_input', dt_k1_tex_2)).
fof(d3_tex_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>(v1_tex_2(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>v1_subset_1(X3,u1_struct_0(X1))))))), file('data/MPTP/proofs/0144bb96470949329693c24e483a6486.E_input', d3_tex_2)).
fof(t1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/0144bb96470949329693c24e483a6486.E_input', t1_tsep_1)).
fof(t6_tex_2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,X1)=>~((v1_subset_1(k6_domain_1(X1,X2),X1)&v1_zfmisc_1(X1))))), file('data/MPTP/proofs/0144bb96470949329693c24e483a6486.E_input', t6_tex_2)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/0144bb96470949329693c24e483a6486.E_input', dt_m1_pre_topc)).
fof(cc2_tex_2, axiom, ![X1]:((~(v1_xboole_0(X1))&v1_zfmisc_1(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(~(v1_xboole_0(X2))=>~(v1_subset_1(X2,X1))))), file('data/MPTP/proofs/0144bb96470949329693c24e483a6486.E_input', cc2_tex_2)).
fof(t7_tex_2, axiom, ![X1]:((~(v1_xboole_0(X1))&~(v1_zfmisc_1(X1)))=>![X2]:(m1_subset_1(X2,X1)=>v1_subset_1(k6_domain_1(X1,X2),X1))), file('data/MPTP/proofs/0144bb96470949329693c24e483a6486.E_input', t7_tex_2)).
fof(d7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(v1_subset_1(X2,X1)<=>~(X2=X1))), file('data/MPTP/proofs/0144bb96470949329693c24e483a6486.E_input', d7_subset_1)).
fof(t8_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>~((v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1))&v7_struct_0(X1))))), file('data/MPTP/proofs/0144bb96470949329693c24e483a6486.E_input', t8_tex_2)).
fof(fc1_struct_0, axiom, ![X1]:((v2_struct_0(X1)&l1_struct_0(X1))=>v1_xboole_0(u1_struct_0(X1))), file('data/MPTP/proofs/0144bb96470949329693c24e483a6486.E_input', fc1_struct_0)).
fof(fc2_tex_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>(~(v2_struct_0(k1_tex_2(X1,X2)))&(v7_struct_0(k1_tex_2(X1,X2))&v1_pre_topc(k1_tex_2(X1,X2))))), file('data/MPTP/proofs/0144bb96470949329693c24e483a6486.E_input', fc2_tex_2)).
fof(c_0_14, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_tex_2(k1_tex_2(X1,X2),X1)<=>v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1)))))), inference(assume_negation,[status(cth)],[t20_tex_2])).
fof(c_0_15, plain, ![X6]:(~l1_pre_topc(X6)|l1_struct_0(X6)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_16, negated_conjecture, ((~v2_struct_0(esk1_0)&l1_pre_topc(esk1_0))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&((~v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)|~v1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),u1_struct_0(esk1_0)))&(v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)|v1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),u1_struct_0(esk1_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])).
fof(c_0_17, plain, ![X7]:(v2_struct_0(X7)|~l1_struct_0(X7)|~v1_xboole_0(u1_struct_0(X7))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_18, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_20, plain, ![X8, X9]:((~v2_struct_0(k1_tex_2(X8,X9))|(v2_struct_0(X8)|~l1_pre_topc(X8)|~m1_subset_1(X9,u1_struct_0(X8))))&((v1_pre_topc(k1_tex_2(X8,X9))|(v2_struct_0(X8)|~l1_pre_topc(X8)|~m1_subset_1(X9,u1_struct_0(X8))))&(m1_pre_topc(k1_tex_2(X8,X9),X8)|(v2_struct_0(X8)|~l1_pre_topc(X8)|~m1_subset_1(X9,u1_struct_0(X8)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_tex_2])])])])).
fof(c_0_21, plain, ![X10, X11, X12]:((~v1_tex_2(X11,X10)|(~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X10)))|(X12!=u1_struct_0(X11)|v1_subset_1(X12,u1_struct_0(X10))))|~m1_pre_topc(X11,X10)|~l1_pre_topc(X10))&((m1_subset_1(esk3_2(X10,X11),k1_zfmisc_1(u1_struct_0(X10)))|v1_tex_2(X11,X10)|~m1_pre_topc(X11,X10)|~l1_pre_topc(X10))&((esk3_2(X10,X11)=u1_struct_0(X11)|v1_tex_2(X11,X10)|~m1_pre_topc(X11,X10)|~l1_pre_topc(X10))&(~v1_subset_1(esk3_2(X10,X11),u1_struct_0(X10))|v1_tex_2(X11,X10)|~m1_pre_topc(X11,X10)|~l1_pre_topc(X10))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tex_2])])])])])).
fof(c_0_22, plain, ![X14, X15]:(~l1_pre_topc(X14)|(~m1_pre_topc(X15,X14)|m1_subset_1(u1_struct_0(X15),k1_zfmisc_1(u1_struct_0(X14))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])).
fof(c_0_23, plain, ![X16, X17]:(v1_xboole_0(X16)|(~m1_subset_1(X17,X16)|(~v1_subset_1(k6_domain_1(X16,X17),X16)|~v1_zfmisc_1(X16)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_tex_2])])])])).
cnf(c_0_24, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, negated_conjecture, (l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
fof(c_0_27, plain, ![X18, X19]:(~l1_pre_topc(X18)|(~m1_pre_topc(X19,X18)|l1_pre_topc(X19))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_28, plain, (m1_pre_topc(k1_tex_2(X1,X2),X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_30, plain, ![X20, X21]:(v1_xboole_0(X20)|~v1_zfmisc_1(X20)|(~m1_subset_1(X21,k1_zfmisc_1(X20))|(v1_xboole_0(X21)|~v1_subset_1(X21,X20)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_tex_2])])])])).
cnf(c_0_31, plain, (v1_subset_1(X3,u1_struct_0(X2))|~v1_tex_2(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|X3!=u1_struct_0(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_32, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33, plain, (v1_xboole_0(X1)|~m1_subset_1(X2,X1)|~v1_subset_1(k6_domain_1(X1,X2),X1)|~v1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_34, negated_conjecture, (v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)|v1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_35, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26])])).
cnf(c_0_36, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, negated_conjecture, (m1_pre_topc(k1_tex_2(esk1_0,esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_19])]), c_0_24])).
fof(c_0_38, plain, ![X22, X23]:(v1_xboole_0(X22)|v1_zfmisc_1(X22)|(~m1_subset_1(X23,X22)|v1_subset_1(k6_domain_1(X22,X23),X22))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_tex_2])])])])).
cnf(c_0_39, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_zfmisc_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~v1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40, plain, (v1_subset_1(u1_struct_0(X1),u1_struct_0(X2))|u1_struct_0(X1)!=u1_struct_0(X3)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)|~v1_tex_2(X3,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_41, negated_conjecture, (v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)|~v1_zfmisc_1(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_29])]), c_0_35])).
cnf(c_0_42, plain, (v2_struct_0(X1)|~v2_struct_0(k1_tex_2(X1,X2))|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_43, negated_conjecture, (l1_pre_topc(k1_tex_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_19])])).
fof(c_0_44, plain, ![X24, X25]:((~v1_subset_1(X25,X24)|X25!=X24|~m1_subset_1(X25,k1_zfmisc_1(X24)))&(X25=X24|v1_subset_1(X25,X24)|~m1_subset_1(X25,k1_zfmisc_1(X24)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_subset_1])])])])).
cnf(c_0_45, negated_conjecture, (~v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)|~v1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_46, plain, (v1_xboole_0(X1)|v1_zfmisc_1(X1)|v1_subset_1(k6_domain_1(X1,X2),X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_47, plain, (esk3_2(X1,X2)=u1_struct_0(X2)|v1_tex_2(X2,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_48, plain, (v1_xboole_0(u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X2))|~v1_zfmisc_1(u1_struct_0(X1))|~m1_pre_topc(X2,X1)|~v1_subset_1(u1_struct_0(X2),u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_39, c_0_32])).
cnf(c_0_49, negated_conjecture, (v1_subset_1(u1_struct_0(X1),u1_struct_0(esk1_0))|u1_struct_0(X1)!=u1_struct_0(k1_tex_2(esk1_0,esk2_0))|~v1_zfmisc_1(u1_struct_0(esk1_0))|~m1_pre_topc(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_37]), c_0_19])])).
cnf(c_0_50, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(k1_tex_2(X1,X2)))|~l1_struct_0(k1_tex_2(X1,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_42, c_0_25])).
cnf(c_0_51, negated_conjecture, (l1_struct_0(k1_tex_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_18, c_0_43])).
cnf(c_0_52, plain, (X1=X2|v1_subset_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_53, plain, (m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|v1_tex_2(X2,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_54, plain, (v1_tex_2(X2,X1)|~v1_subset_1(esk3_2(X1,X2),u1_struct_0(X1))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_55, negated_conjecture, (v1_zfmisc_1(u1_struct_0(esk1_0))|~v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_29])]), c_0_35])).
cnf(c_0_56, negated_conjecture, (u1_struct_0(X1)=esk3_2(esk1_0,X1)|v1_tex_2(X1,esk1_0)|~m1_pre_topc(X1,esk1_0)), inference(spm,[status(thm)],[c_0_47, c_0_19])).
cnf(c_0_57, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|u1_struct_0(X1)!=u1_struct_0(k1_tex_2(esk1_0,esk2_0))|~v1_zfmisc_1(u1_struct_0(esk1_0))|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_19])]), c_0_35])).
cnf(c_0_58, negated_conjecture, (~v1_xboole_0(u1_struct_0(k1_tex_2(esk1_0,esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_29]), c_0_51]), c_0_19])]), c_0_24])).
cnf(c_0_59, plain, (esk3_2(X1,X2)=u1_struct_0(X1)|v1_tex_2(X2,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54])).
fof(c_0_60, plain, ![X26, X27]:(v2_struct_0(X26)|~l1_struct_0(X26)|(~m1_subset_1(X27,u1_struct_0(X26))|(~v1_subset_1(k6_domain_1(u1_struct_0(X26),X27),u1_struct_0(X26))|~v7_struct_0(X26)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_tex_2])])])])).
fof(c_0_61, plain, ![X28]:(~v2_struct_0(X28)|~l1_struct_0(X28)|v1_xboole_0(u1_struct_0(X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_struct_0])])).
cnf(c_0_62, negated_conjecture, (u1_struct_0(k1_tex_2(esk1_0,esk2_0))=esk3_2(esk1_0,k1_tex_2(esk1_0,esk2_0))|v1_zfmisc_1(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_37])])).
cnf(c_0_63, negated_conjecture, (~v1_zfmisc_1(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_57]), c_0_37])]), c_0_58])).
cnf(c_0_64, negated_conjecture, (esk3_2(esk1_0,X1)=u1_struct_0(esk1_0)|v1_tex_2(X1,esk1_0)|~m1_pre_topc(X1,esk1_0)), inference(spm,[status(thm)],[c_0_59, c_0_19])).
fof(c_0_65, plain, ![X29, X30]:((~v2_struct_0(k1_tex_2(X29,X30))|(v2_struct_0(X29)|~l1_pre_topc(X29)|~m1_subset_1(X30,u1_struct_0(X29))))&((v7_struct_0(k1_tex_2(X29,X30))|(v2_struct_0(X29)|~l1_pre_topc(X29)|~m1_subset_1(X30,u1_struct_0(X29))))&(v1_pre_topc(k1_tex_2(X29,X30))|(v2_struct_0(X29)|~l1_pre_topc(X29)|~m1_subset_1(X30,u1_struct_0(X29)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_tex_2])])])])).
cnf(c_0_66, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1))|~v7_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_67, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_68, negated_conjecture, (u1_struct_0(k1_tex_2(esk1_0,esk2_0))=esk3_2(esk1_0,k1_tex_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_69, negated_conjecture, (esk3_2(esk1_0,k1_tex_2(esk1_0,esk2_0))=u1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_64]), c_0_37])]), c_0_63])).
cnf(c_0_70, plain, (v7_struct_0(k1_tex_2(X1,X2))|v2_struct_0(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_71, plain, (v1_zfmisc_1(u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|~v7_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_46]), c_0_67])).
cnf(c_0_72, negated_conjecture, (u1_struct_0(k1_tex_2(esk1_0,esk2_0))=u1_struct_0(esk1_0)), inference(rw,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_73, negated_conjecture, (v7_struct_0(k1_tex_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_29]), c_0_19])]), c_0_24])).
cnf(c_0_74, negated_conjecture, (~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73]), c_0_51])]), c_0_63]), c_0_35])).
cnf(c_0_75, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_29, c_0_74]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 76
# Proof object clause steps            : 47
# Proof object formula steps           : 29
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 22
# Proof object simplifying inferences  : 44
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 26
# Processed clauses                    : 92
# ...of these trivial                  : 0
# ...subsumed                          : 5
# ...remaining for further processing  : 87
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 11
# Generated clauses                    : 108
# ...of the previous two non-trivial   : 110
# Contextual simplify-reflections      : 2
# Paramodulations                      : 104
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 49
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 35
# Current number of unprocessed clauses: 68
# ...number of literals in the above   : 258
# Current number of archived formulas  : 0
# Current number of archived clauses   : 37
# Clause-clause subsumption calls (NU) : 459
# Rec. Clause-clause subsumption calls : 197
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 40
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4604

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28616 pages
