# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t15_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4))))))), file('data/MPTP/proofs/018d50a9b493478586c65eaebb1443df.E_input', t15_yellow_0)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/018d50a9b493478586c65eaebb1443df.E_input', d2_subset_1)).
fof(d5_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>(v2_yellow_0(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&r2_lattice3(X1,u1_struct_0(X1),X2)))), file('data/MPTP/proofs/018d50a9b493478586c65eaebb1443df.E_input', d5_yellow_0)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/018d50a9b493478586c65eaebb1443df.E_input', d9_lattice3)).
fof(t8_boole, axiom, ![X1, X2]:~((v1_xboole_0(X1)&(~(X1=X2)&v1_xboole_0(X2)))), file('data/MPTP/proofs/018d50a9b493478586c65eaebb1443df.E_input', t8_boole)).
fof(t16_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3))))))), file('data/MPTP/proofs/018d50a9b493478586c65eaebb1443df.E_input', t16_yellow_0)).
fof(t6_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_lattice3(X1,k1_xboole_0,X2)&r1_lattice3(X1,k1_xboole_0,X2)))), file('data/MPTP/proofs/018d50a9b493478586c65eaebb1443df.E_input', t6_yellow_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/018d50a9b493478586c65eaebb1443df.E_input', fc1_xboole_0)).
fof(t43_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v2_yellow_0(X1)&l1_orders_2(X1))))=>(r2_yellow_0(X1,k1_xboole_0)&r1_yellow_0(X1,u1_struct_0(X1)))), file('data/MPTP/proofs/018d50a9b493478586c65eaebb1443df.E_input', t43_yellow_0)).
fof(d4_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>(v1_yellow_0(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&r1_lattice3(X1,u1_struct_0(X1),X2)))), file('data/MPTP/proofs/018d50a9b493478586c65eaebb1443df.E_input', d4_yellow_0)).
fof(t42_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))=>(r1_yellow_0(X1,k1_xboole_0)&r2_yellow_0(X1,u1_struct_0(X1)))), file('data/MPTP/proofs/018d50a9b493478586c65eaebb1443df.E_input', t42_yellow_0)).
fof(c_0_11, plain, ![X21, X22, X24, X25, X26]:(((m1_subset_1(esk4_2(X21,X22),u1_struct_0(X21))|~r1_yellow_0(X21,X22)|(~v5_orders_2(X21)|~l1_orders_2(X21)))&((r2_lattice3(X21,X22,esk4_2(X21,X22))|~r1_yellow_0(X21,X22)|(~v5_orders_2(X21)|~l1_orders_2(X21)))&(~m1_subset_1(X24,u1_struct_0(X21))|(~r2_lattice3(X21,X22,X24)|r1_orders_2(X21,esk4_2(X21,X22),X24))|~r1_yellow_0(X21,X22)|(~v5_orders_2(X21)|~l1_orders_2(X21)))))&((m1_subset_1(esk5_3(X21,X25,X26),u1_struct_0(X21))|~r2_lattice3(X21,X25,X26)|~m1_subset_1(X26,u1_struct_0(X21))|r1_yellow_0(X21,X25)|(~v5_orders_2(X21)|~l1_orders_2(X21)))&((r2_lattice3(X21,X25,esk5_3(X21,X25,X26))|~r2_lattice3(X21,X25,X26)|~m1_subset_1(X26,u1_struct_0(X21))|r1_yellow_0(X21,X25)|(~v5_orders_2(X21)|~l1_orders_2(X21)))&(~r1_orders_2(X21,X26,esk5_3(X21,X25,X26))|~r2_lattice3(X21,X25,X26)|~m1_subset_1(X26,u1_struct_0(X21))|r1_yellow_0(X21,X25)|(~v5_orders_2(X21)|~l1_orders_2(X21)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_yellow_0])])])])])])).
fof(c_0_12, plain, ![X7, X8]:(((~m1_subset_1(X8,X7)|r2_hidden(X8,X7)|v1_xboole_0(X7))&(~r2_hidden(X8,X7)|m1_subset_1(X8,X7)|v1_xboole_0(X7)))&((~m1_subset_1(X8,X7)|v1_xboole_0(X8)|~v1_xboole_0(X7))&(~v1_xboole_0(X8)|m1_subset_1(X8,X7)|~v1_xboole_0(X7)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_13, plain, ![X28, X30]:(((m1_subset_1(esk6_1(X28),u1_struct_0(X28))|~v2_yellow_0(X28)|~l1_orders_2(X28))&(r2_lattice3(X28,u1_struct_0(X28),esk6_1(X28))|~v2_yellow_0(X28)|~l1_orders_2(X28)))&(~m1_subset_1(X30,u1_struct_0(X28))|~r2_lattice3(X28,u1_struct_0(X28),X30)|v2_yellow_0(X28)|~l1_orders_2(X28))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_yellow_0])])])])])).
fof(c_0_14, plain, ![X16, X17, X18, X19]:((~r2_lattice3(X16,X17,X18)|(~m1_subset_1(X19,u1_struct_0(X16))|(~r2_hidden(X19,X17)|r1_orders_2(X16,X19,X18)))|~m1_subset_1(X18,u1_struct_0(X16))|~l1_orders_2(X16))&((m1_subset_1(esk3_3(X16,X17,X18),u1_struct_0(X16))|r2_lattice3(X16,X17,X18)|~m1_subset_1(X18,u1_struct_0(X16))|~l1_orders_2(X16))&((r2_hidden(esk3_3(X16,X17,X18),X17)|r2_lattice3(X16,X17,X18)|~m1_subset_1(X18,u1_struct_0(X16))|~l1_orders_2(X16))&(~r1_orders_2(X16,esk3_3(X16,X17,X18),X18)|r2_lattice3(X16,X17,X18)|~m1_subset_1(X18,u1_struct_0(X16))|~l1_orders_2(X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
fof(c_0_15, plain, ![X5, X6]:(~v1_xboole_0(X5)|(X5=X6|~v1_xboole_0(X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_boole])])])).
fof(c_0_16, plain, ![X9, X10, X12, X13, X14]:(((m1_subset_1(esk1_2(X9,X10),u1_struct_0(X9))|~r2_yellow_0(X9,X10)|(~v5_orders_2(X9)|~l1_orders_2(X9)))&((r1_lattice3(X9,X10,esk1_2(X9,X10))|~r2_yellow_0(X9,X10)|(~v5_orders_2(X9)|~l1_orders_2(X9)))&(~m1_subset_1(X12,u1_struct_0(X9))|(~r1_lattice3(X9,X10,X12)|r1_orders_2(X9,X12,esk1_2(X9,X10)))|~r2_yellow_0(X9,X10)|(~v5_orders_2(X9)|~l1_orders_2(X9)))))&((m1_subset_1(esk2_3(X9,X13,X14),u1_struct_0(X9))|~r1_lattice3(X9,X13,X14)|~m1_subset_1(X14,u1_struct_0(X9))|r2_yellow_0(X9,X13)|(~v5_orders_2(X9)|~l1_orders_2(X9)))&((r1_lattice3(X9,X13,esk2_3(X9,X13,X14))|~r1_lattice3(X9,X13,X14)|~m1_subset_1(X14,u1_struct_0(X9))|r2_yellow_0(X9,X13)|(~v5_orders_2(X9)|~l1_orders_2(X9)))&(~r1_orders_2(X9,esk2_3(X9,X13,X14),X14)|~r1_lattice3(X9,X13,X14)|~m1_subset_1(X14,u1_struct_0(X9))|r2_yellow_0(X9,X13)|(~v5_orders_2(X9)|~l1_orders_2(X9)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_yellow_0])])])])])])).
cnf(c_0_17, plain, (r1_yellow_0(X1,X3)|~r1_orders_2(X1,X2,esk5_3(X1,X3,X2))|~r2_lattice3(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, plain, (r1_orders_2(X2,esk4_2(X2,X3),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r1_yellow_0(X2,X3)|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19, plain, (m1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_20, plain, ![X31, X32]:((r2_lattice3(X31,k1_xboole_0,X32)|~m1_subset_1(X32,u1_struct_0(X31))|~l1_orders_2(X31))&(r1_lattice3(X31,k1_xboole_0,X32)|~m1_subset_1(X32,u1_struct_0(X31))|~l1_orders_2(X31))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_yellow_0])])])])).
cnf(c_0_21, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22, plain, (m1_subset_1(esk6_1(X1),u1_struct_0(X1))|~v2_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23, plain, (r1_orders_2(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_24, plain, (r2_lattice3(X1,X2,esk5_3(X1,X2,X3))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_25, plain, (m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_26, plain, (X1=X2|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_27, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_28, plain, (m1_subset_1(esk1_2(X1,X2),u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_29, plain, (r1_yellow_0(X1,X2)|~r1_yellow_0(X1,X3)|~r2_lattice3(X1,X3,esk5_3(X1,X2,esk4_2(X1,X3)))|~r2_lattice3(X1,X2,esk4_2(X1,X3))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(esk5_3(X1,X2,esk4_2(X1,X3)),u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])).
cnf(c_0_30, plain, (r2_lattice3(X1,k1_xboole_0,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_31, plain, (v1_xboole_0(esk6_1(X1))|~v2_yellow_0(X1)|~l1_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_32, plain, (r1_yellow_0(X1,X2)|r1_orders_2(X1,X3,esk5_3(X1,X2,X4))|~r2_lattice3(X1,X2,X4)|~l1_orders_2(X1)|~v5_orders_2(X1)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
cnf(c_0_33, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_34, plain, (v1_xboole_0(esk1_2(X1,X2))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_21, c_0_28])).
cnf(c_0_35, plain, (m1_subset_1(esk3_3(X1,X2,X3),u1_struct_0(X1))|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_36, plain, (r1_yellow_0(X1,X2)|~r1_yellow_0(X1,k1_xboole_0)|~r2_lattice3(X1,X2,esk4_2(X1,k1_xboole_0))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(esk5_3(X1,X2,esk4_2(X1,k1_xboole_0)),u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_37, plain, (X1=esk6_1(X2)|~v2_yellow_0(X2)|~l1_orders_2(X2)|~v1_xboole_0(u1_struct_0(X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_26, c_0_31])).
cnf(c_0_38, plain, (v1_xboole_0(esk4_2(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_21, c_0_19])).
cnf(c_0_39, plain, (r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_17, c_0_32])).
cnf(c_0_40, plain, (r2_lattice3(X1,u1_struct_0(X1),esk6_1(X1))|~v2_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_41, plain, (r1_yellow_0(X1,X2)|v1_xboole_0(esk5_3(X1,X2,X3))|~r2_lattice3(X1,X2,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_21, c_0_25])).
cnf(c_0_42, plain, (r1_orders_2(X2,X1,esk1_2(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~r2_yellow_0(X2,X3)|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_43, plain, (esk1_2(X1,X2)=k1_xboole_0|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_44, plain, (r2_lattice3(X1,X2,X3)|v1_xboole_0(esk3_3(X1,X2,X3))|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_21, c_0_35])).
cnf(c_0_45, plain, (r1_yellow_0(X1,X2)|~r1_yellow_0(X1,k1_xboole_0)|~r2_lattice3(X1,X2,esk4_2(X1,k1_xboole_0))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_25]), c_0_19])).
cnf(c_0_46, plain, (esk4_2(X1,X2)=esk6_1(X3)|~v2_yellow_0(X3)|~r1_yellow_0(X1,X2)|~l1_orders_2(X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v1_xboole_0(u1_struct_0(X3))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_47, plain, (r1_yellow_0(X1,u1_struct_0(X1))|~v2_yellow_0(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~r2_hidden(esk6_1(X1),u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_22])).
cnf(c_0_48, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_49, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v2_yellow_0(X1)&l1_orders_2(X1))))=>(r2_yellow_0(X1,k1_xboole_0)&r1_yellow_0(X1,u1_struct_0(X1))))), inference(assume_negation,[status(cth)],[t43_yellow_0])).
cnf(c_0_50, plain, (esk5_3(X1,X2,X3)=k1_xboole_0|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_33, c_0_41])).
cnf(c_0_51, plain, (r1_orders_2(X1,X2,k1_xboole_0)|~r1_lattice3(X1,X3,X2)|~r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_52, plain, (r1_lattice3(X1,k1_xboole_0,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_53, plain, (r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk3_3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_54, plain, (esk3_3(X1,X2,X3)=k1_xboole_0|r2_lattice3(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_33, c_0_44])).
cnf(c_0_55, plain, (r1_lattice3(X1,X2,esk1_2(X1,X2))|~r2_yellow_0(X1,X2)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_56, plain, (r2_yellow_0(X1,X2)|~r1_orders_2(X1,esk2_3(X1,X2,X3),X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_57, plain, (r1_orders_2(X1,X2,esk6_1(X1))|~v2_yellow_0(X1)|~l1_orders_2(X1)|~r2_hidden(X2,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_40]), c_0_22])).
cnf(c_0_58, plain, (r1_yellow_0(X1,X2)|~v2_yellow_0(X3)|~r1_yellow_0(X1,k1_xboole_0)|~r2_lattice3(X1,X2,esk6_1(X3))|~l1_orders_2(X1)|~l1_orders_2(X3)|~v5_orders_2(X1)|~v1_xboole_0(u1_struct_0(X3))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_59, plain, (r1_yellow_0(X1,u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|~v2_yellow_0(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_22])).
fof(c_0_60, negated_conjecture, ((~v2_struct_0(esk7_0)&(v5_orders_2(esk7_0)&(v2_yellow_0(esk7_0)&l1_orders_2(esk7_0))))&(~r2_yellow_0(esk7_0,k1_xboole_0)|~r1_yellow_0(esk7_0,u1_struct_0(esk7_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_49])])])])).
cnf(c_0_61, plain, (r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,X3,k1_xboole_0)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_17, c_0_50])).
cnf(c_0_62, plain, (r1_orders_2(X1,X2,k1_xboole_0)|~r2_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_63, plain, (r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,k1_xboole_0,X3)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_64, plain, (m1_subset_1(k1_xboole_0,u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_28, c_0_43])).
cnf(c_0_65, plain, (r1_lattice3(X1,X2,k1_xboole_0)|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_55, c_0_43])).
cnf(c_0_66, plain, (m1_subset_1(esk2_3(X1,X2,X3),u1_struct_0(X1))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_67, plain, (r2_yellow_0(X1,X2)|~v2_yellow_0(X1)|~r1_lattice3(X1,X2,esk6_1(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~r2_hidden(esk2_3(X1,X2,esk6_1(X1)),u1_struct_0(X1))|~m1_subset_1(esk2_3(X1,X2,esk6_1(X1)),u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_22])).
cnf(c_0_68, plain, (r1_yellow_0(X1,u1_struct_0(X1))|~v2_yellow_0(X1)|~r1_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_40]), c_0_59])).
cnf(c_0_69, negated_conjecture, (v2_yellow_0(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_70, negated_conjecture, (l1_orders_2(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_71, negated_conjecture, (v5_orders_2(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_72, plain, (r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~r2_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_73, plain, (v2_yellow_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,u1_struct_0(X2),X1)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_74, plain, (r2_lattice3(X1,X2,esk1_2(X1,X3))|~r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_42]), c_0_64]), c_0_28]), c_0_65])).
cnf(c_0_75, plain, (r2_yellow_0(X1,X2)|v1_xboole_0(esk2_3(X1,X2,X3))|~r1_lattice3(X1,X2,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_21, c_0_66])).
cnf(c_0_76, plain, (esk4_2(X1,X2)=k1_xboole_0|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_33, c_0_38])).
cnf(c_0_77, plain, (r2_lattice3(X1,X2,esk4_2(X1,X2))|~r1_yellow_0(X1,X2)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_78, plain, (r2_yellow_0(X1,X2)|v1_xboole_0(u1_struct_0(X1))|~v2_yellow_0(X1)|~r1_lattice3(X1,X2,esk6_1(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(esk2_3(X1,X2,esk6_1(X1)),u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_67, c_0_48])).
cnf(c_0_79, negated_conjecture, (~r2_yellow_0(esk7_0,k1_xboole_0)|~r1_yellow_0(esk7_0,u1_struct_0(esk7_0))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_80, negated_conjecture, (r1_yellow_0(esk7_0,u1_struct_0(esk7_0))|~r1_yellow_0(esk7_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70]), c_0_71])])).
cnf(c_0_81, plain, (r1_yellow_0(X1,k1_xboole_0)|~r2_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_72, c_0_30])).
cnf(c_0_82, plain, (v2_yellow_0(X1)|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_28])).
cnf(c_0_83, negated_conjecture, (r1_yellow_0(esk7_0,u1_struct_0(esk7_0))|v1_xboole_0(u1_struct_0(esk7_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_69]), c_0_70]), c_0_71])])).
cnf(c_0_84, plain, (r1_lattice3(X1,X2,esk2_3(X1,X2,X3))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_85, plain, (esk2_3(X1,X2,X3)=k1_xboole_0|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_33, c_0_75])).
cnf(c_0_86, plain, (r1_orders_2(X1,k1_xboole_0,X2)|~r1_yellow_0(X1,X3)|~r2_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_18, c_0_76])).
cnf(c_0_87, plain, (v2_yellow_0(X1)|~r1_yellow_0(X1,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_77]), c_0_19])).
cnf(c_0_88, plain, (r2_yellow_0(X1,X2)|v1_xboole_0(u1_struct_0(X1))|~v2_yellow_0(X1)|~r1_lattice3(X1,X2,esk6_1(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_66]), c_0_22])).
cnf(c_0_89, negated_conjecture, (~r1_yellow_0(esk7_0,k1_xboole_0)|~r2_yellow_0(esk7_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_90, plain, (r1_yellow_0(X1,k1_xboole_0)|~r2_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_22]), c_0_82])).
cnf(c_0_91, negated_conjecture, (v1_xboole_0(u1_struct_0(esk7_0))|~r2_yellow_0(esk7_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_79, c_0_83])).
cnf(c_0_92, plain, (r1_lattice3(X1,X2,k1_xboole_0)|~r1_lattice3(X1,X2,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_65])).
cnf(c_0_93, plain, (r1_orders_2(X1,k1_xboole_0,esk6_1(X1))|~r1_yellow_0(X1,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_40]), c_0_22]), c_0_87])).
cnf(c_0_94, plain, (r2_yellow_0(X1,k1_xboole_0)|v1_xboole_0(u1_struct_0(X1))|~v2_yellow_0(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_52]), c_0_22])).
cnf(c_0_95, negated_conjecture, (~r2_yellow_0(esk7_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_70]), c_0_71])]), c_0_91])).
cnf(c_0_96, plain, (r1_lattice3(X1,k1_xboole_0,k1_xboole_0)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_92, c_0_52])).
cnf(c_0_97, plain, (r2_yellow_0(X1,X2)|~r1_orders_2(X1,k1_xboole_0,X3)|~r1_lattice3(X1,X2,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_56, c_0_85])).
cnf(c_0_98, plain, (r2_lattice3(X1,X2,esk6_1(X1))|~r1_yellow_0(X1,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(esk6_1(X1),u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_63, c_0_93])).
fof(c_0_99, plain, ![X34, X36]:(((m1_subset_1(esk8_1(X34),u1_struct_0(X34))|~v1_yellow_0(X34)|~l1_orders_2(X34))&(r1_lattice3(X34,u1_struct_0(X34),esk8_1(X34))|~v1_yellow_0(X34)|~l1_orders_2(X34)))&(~m1_subset_1(X36,u1_struct_0(X34))|~r1_lattice3(X34,u1_struct_0(X34),X36)|v1_yellow_0(X34)|~l1_orders_2(X34))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_yellow_0])])])])])).
cnf(c_0_100, negated_conjecture, (v1_xboole_0(u1_struct_0(esk7_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_69]), c_0_70]), c_0_71])]), c_0_95])).
cnf(c_0_101, plain, (r1_lattice3(X1,k1_xboole_0,k1_xboole_0)|~v2_yellow_0(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_96, c_0_22])).
cnf(c_0_102, plain, (m1_subset_1(X1,X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_103, plain, (r2_yellow_0(X1,X2)|~r1_yellow_0(X1,u1_struct_0(X1))|~r1_lattice3(X1,X2,esk6_1(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(esk6_1(X1),u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_97, c_0_93])).
cnf(c_0_104, plain, (m1_subset_1(esk6_1(X1),u1_struct_0(X2))|~v2_yellow_0(X1)|~r1_yellow_0(X2,X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v5_orders_2(X2)|~v1_xboole_0(u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X2))), inference(spm,[status(thm)],[c_0_19, c_0_46])).
cnf(c_0_105, plain, (r1_yellow_0(X1,X2)|~v2_yellow_0(X1)|~r1_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_98]), c_0_22]), c_0_68])).
fof(c_0_106, plain, ![X37]:((r1_yellow_0(X37,k1_xboole_0)|(v2_struct_0(X37)|(~v5_orders_2(X37)|(~v1_yellow_0(X37)|~l1_orders_2(X37)))))&(r2_yellow_0(X37,u1_struct_0(X37))|(v2_struct_0(X37)|(~v5_orders_2(X37)|(~v1_yellow_0(X37)|~l1_orders_2(X37)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t42_yellow_0])])])])).
cnf(c_0_107, plain, (v1_yellow_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,u1_struct_0(X2),X1)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_108, negated_conjecture, (u1_struct_0(esk7_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_33, c_0_100])).
cnf(c_0_109, negated_conjecture, (r1_lattice3(esk7_0,k1_xboole_0,k1_xboole_0)|~v1_xboole_0(u1_struct_0(esk7_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_69]), c_0_70]), c_0_71])])).
cnf(c_0_110, plain, (m1_subset_1(X1,k1_xboole_0)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_102, c_0_27])).
cnf(c_0_111, plain, (r2_yellow_0(X1,k1_xboole_0)|~r1_yellow_0(X1,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(esk6_1(X1),u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_103, c_0_52])).
cnf(c_0_112, negated_conjecture, (m1_subset_1(esk6_1(X1),u1_struct_0(esk7_0))|~v2_yellow_0(X1)|~r1_yellow_0(esk7_0,k1_xboole_0)|~l1_orders_2(X1)|~v1_xboole_0(u1_struct_0(esk7_0))|~v1_xboole_0(u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_80]), c_0_70]), c_0_71])])).
cnf(c_0_113, negated_conjecture, (r1_yellow_0(esk7_0,X1)|~r1_yellow_0(esk7_0,k1_xboole_0)|~v1_xboole_0(u1_struct_0(esk7_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_69]), c_0_70]), c_0_71])])).
cnf(c_0_114, negated_conjecture, (~v2_struct_0(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_115, plain, (r1_yellow_0(X1,k1_xboole_0)|v2_struct_0(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_116, negated_conjecture, (v1_yellow_0(esk7_0)|~r1_lattice3(esk7_0,k1_xboole_0,X1)|~m1_subset_1(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_70])])).
cnf(c_0_117, negated_conjecture, (r1_lattice3(esk7_0,k1_xboole_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_109, c_0_100])])).
cnf(c_0_118, plain, (m1_subset_1(k1_xboole_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_110, c_0_27])).
cnf(c_0_119, negated_conjecture, (~r1_yellow_0(esk7_0,k1_xboole_0)|~v1_xboole_0(u1_struct_0(esk7_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_70]), c_0_71]), c_0_69])]), c_0_113]), c_0_89])).
cnf(c_0_120, negated_conjecture, (r1_yellow_0(esk7_0,k1_xboole_0)|~v1_yellow_0(esk7_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_70]), c_0_71])])).
cnf(c_0_121, negated_conjecture, (v1_yellow_0(esk7_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_118])])).
cnf(c_0_122, negated_conjecture, (~r1_yellow_0(esk7_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_119, c_0_100])])).
cnf(c_0_123, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_120, c_0_121])]), c_0_122]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 124
# Proof object clause steps            : 102
# Proof object formula steps           : 22
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 32
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 67
# Proof object simplifying inferences  : 64
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 37
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 37
# Processed clauses                    : 1162
# ...of these trivial                  : 3
# ...subsumed                          : 666
# ...remaining for further processing  : 493
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 95
# Backward-rewritten                   : 31
# Generated clauses                    : 3302
# ...of the previous two non-trivial   : 2911
# Contextual simplify-reflections      : 120
# Paramodulations                      : 3302
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 330
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 316
# Current number of unprocessed clauses: 1675
# ...number of literals in the above   : 18760
# Current number of archived formulas  : 0
# Current number of archived clauses   : 163
# Clause-clause subsumption calls (NU) : 116787
# Rec. Clause-clause subsumption calls : 2949
# Non-unit clause-clause subsumptions  : 831
# Unit Clause-clause subsumption calls : 1115
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 95093

# -------------------------------------------------
# User time                : 0.084 s
# System time              : 0.003 s
# Total time               : 0.087 s
# Maximum resident set size: 28876 pages
