# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t166_funct_1, conjecture, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((k1_relat_1(X1)=k1_relat_1(X2)&![X4]:(r2_hidden(X4,X3)=>k1_funct_1(X1,X4)=k1_funct_1(X2,X4)))=>k7_relat_1(X1,X3)=k7_relat_1(X2,X3)))), file('data/MPTP/proofs/01c715b94774422da9888eb5e4c672e1.E_input', t166_funct_1)).
fof(t90_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k1_relat_1(k7_relat_1(X2,X1))=k3_xboole_0(k1_relat_1(X2),X1)), file('data/MPTP/proofs/01c715b94774422da9888eb5e4c672e1.E_input', t90_relat_1)).
fof(t86_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k1_relat_1(k7_relat_1(X3,X2)))<=>(r2_hidden(X1,X2)&r2_hidden(X1,k1_relat_1(X3))))), file('data/MPTP/proofs/01c715b94774422da9888eb5e4c672e1.E_input', t86_relat_1)).
fof(t68_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X2)=k3_xboole_0(k1_relat_1(X3),X1)&![X4]:(r2_hidden(X4,k1_relat_1(X2))=>k1_funct_1(X2,X4)=k1_funct_1(X3,X4)))=>X2=k7_relat_1(X3,X1)))), file('data/MPTP/proofs/01c715b94774422da9888eb5e4c672e1.E_input', t68_funct_1)).
fof(t72_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,X2)=>k1_funct_1(k7_relat_1(X3,X2),X1)=k1_funct_1(X3,X1))), file('data/MPTP/proofs/01c715b94774422da9888eb5e4c672e1.E_input', t72_funct_1)).
fof(fc8_funct_1, axiom, ![X1, X2]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_relat_1(k7_relat_1(X1,X2))&v1_funct_1(k7_relat_1(X1,X2)))), file('data/MPTP/proofs/01c715b94774422da9888eb5e4c672e1.E_input', fc8_funct_1)).
fof(c_0_6, negated_conjecture, ~(![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((k1_relat_1(X1)=k1_relat_1(X2)&![X4]:(r2_hidden(X4,X3)=>k1_funct_1(X1,X4)=k1_funct_1(X2,X4)))=>k7_relat_1(X1,X3)=k7_relat_1(X2,X3))))), inference(assume_negation,[status(cth)],[t166_funct_1])).
fof(c_0_7, plain, ![X21, X22]:(~v1_relat_1(X22)|k1_relat_1(k7_relat_1(X22,X21))=k3_xboole_0(k1_relat_1(X22),X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t90_relat_1])])).
fof(c_0_8, negated_conjecture, ![X8]:((v1_relat_1(esk1_0)&v1_funct_1(esk1_0))&((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&((k1_relat_1(esk1_0)=k1_relat_1(esk2_0)&(~r2_hidden(X8,esk3_0)|k1_funct_1(esk1_0,X8)=k1_funct_1(esk2_0,X8)))&k7_relat_1(esk1_0,esk3_0)!=k7_relat_1(esk2_0,esk3_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
cnf(c_0_9, plain, (k1_relat_1(k7_relat_1(X1,X2))=k3_xboole_0(k1_relat_1(X1),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10, negated_conjecture, (k1_relat_1(esk1_0)=k1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_12, plain, ![X15, X16, X17]:(((r2_hidden(X15,X16)|~r2_hidden(X15,k1_relat_1(k7_relat_1(X17,X16)))|~v1_relat_1(X17))&(r2_hidden(X15,k1_relat_1(X17))|~r2_hidden(X15,k1_relat_1(k7_relat_1(X17,X16)))|~v1_relat_1(X17)))&(~r2_hidden(X15,X16)|~r2_hidden(X15,k1_relat_1(X17))|r2_hidden(X15,k1_relat_1(k7_relat_1(X17,X16)))|~v1_relat_1(X17))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t86_relat_1])])])).
cnf(c_0_13, negated_conjecture, (k3_xboole_0(k1_relat_1(esk2_0),X1)=k1_relat_1(k7_relat_1(esk1_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11])])).
cnf(c_0_14, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_15, plain, ![X9, X10, X11]:((r2_hidden(esk4_3(X9,X10,X11),k1_relat_1(X10))|k1_relat_1(X10)!=k3_xboole_0(k1_relat_1(X11),X9)|X10=k7_relat_1(X11,X9)|(~v1_relat_1(X11)|~v1_funct_1(X11))|(~v1_relat_1(X10)|~v1_funct_1(X10)))&(k1_funct_1(X10,esk4_3(X9,X10,X11))!=k1_funct_1(X11,esk4_3(X9,X10,X11))|k1_relat_1(X10)!=k3_xboole_0(k1_relat_1(X11),X9)|X10=k7_relat_1(X11,X9)|(~v1_relat_1(X11)|~v1_funct_1(X11))|(~v1_relat_1(X10)|~v1_funct_1(X10)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t68_funct_1])])])])])).
fof(c_0_16, plain, ![X18, X19, X20]:(~v1_relat_1(X20)|~v1_funct_1(X20)|(~r2_hidden(X18,X19)|k1_funct_1(k7_relat_1(X20,X19),X18)=k1_funct_1(X20,X18))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t72_funct_1])])).
fof(c_0_17, plain, ![X13, X14]:((v1_relat_1(k7_relat_1(X13,X14))|(~v1_relat_1(X13)|~v1_funct_1(X13)))&(v1_funct_1(k7_relat_1(X13,X14))|(~v1_relat_1(X13)|~v1_funct_1(X13)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc8_funct_1])])])).
cnf(c_0_18, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_relat_1(k7_relat_1(X3,X2)))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, negated_conjecture, (k1_relat_1(k7_relat_1(esk1_0,X1))=k1_relat_1(k7_relat_1(esk2_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_13]), c_0_14])])).
cnf(c_0_20, plain, (X1=k7_relat_1(X3,X2)|k1_funct_1(X1,esk4_3(X2,X1,X3))!=k1_funct_1(X3,esk4_3(X2,X1,X3))|k1_relat_1(X1)!=k3_xboole_0(k1_relat_1(X3),X2)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (k1_funct_1(k7_relat_1(X1,X3),X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (v1_funct_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, negated_conjecture, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_relat_1(k7_relat_1(esk2_0,X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_11])])).
cnf(c_0_25, plain, (r2_hidden(esk4_3(X1,X2,X3),k1_relat_1(X2))|X2=k7_relat_1(X3,X1)|k1_relat_1(X2)!=k3_xboole_0(k1_relat_1(X3),X1)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_26, plain, (k7_relat_1(X1,X2)=k7_relat_1(X3,X4)|k1_funct_1(X1,esk4_3(X4,k7_relat_1(X1,X2),X3))!=k1_funct_1(X3,esk4_3(X4,k7_relat_1(X1,X2),X3))|k1_relat_1(k7_relat_1(X1,X2))!=k3_xboole_0(k1_relat_1(X3),X4)|~r2_hidden(esk4_3(X4,k7_relat_1(X1,X2),X3),X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23])).
cnf(c_0_27, negated_conjecture, (k7_relat_1(esk2_0,X1)=k7_relat_1(X2,X3)|r2_hidden(esk4_3(X3,k7_relat_1(esk2_0,X1),X2),X1)|k1_relat_1(k7_relat_1(esk2_0,X1))!=k3_xboole_0(k1_relat_1(X2),X3)|~v1_funct_1(k7_relat_1(esk2_0,X1))|~v1_funct_1(X2)|~v1_relat_1(k7_relat_1(esk2_0,X1))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_28, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_29, negated_conjecture, (k7_relat_1(esk2_0,X1)=k7_relat_1(X2,X3)|k1_funct_1(esk2_0,esk4_3(X3,k7_relat_1(esk2_0,X1),X2))!=k1_funct_1(X2,esk4_3(X3,k7_relat_1(esk2_0,X1),X2))|k1_relat_1(k7_relat_1(esk2_0,X1))!=k3_xboole_0(k1_relat_1(X2),X3)|~v1_funct_1(k7_relat_1(esk2_0,X1))|~v1_funct_1(X2)|~v1_relat_1(k7_relat_1(esk2_0,X1))|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_14])])).
cnf(c_0_30, negated_conjecture, (k1_funct_1(esk1_0,X1)=k1_funct_1(esk2_0,X1)|~r2_hidden(X1,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_31, negated_conjecture, (v1_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_32, negated_conjecture, (k7_relat_1(esk2_0,X1)=k7_relat_1(esk1_0,X2)|k1_relat_1(k7_relat_1(esk2_0,X1))!=k1_relat_1(k7_relat_1(esk2_0,X2))|~r2_hidden(esk4_3(X2,k7_relat_1(esk2_0,X1),esk1_0),esk3_0)|~v1_funct_1(k7_relat_1(esk2_0,X1))|~v1_relat_1(k7_relat_1(esk2_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_10]), c_0_13]), c_0_19]), c_0_31]), c_0_11])])).
cnf(c_0_33, plain, (k7_relat_1(X1,X2)=k7_relat_1(X3,X4)|r2_hidden(esk4_3(X4,k7_relat_1(X1,X2),X3),X2)|k1_relat_1(k7_relat_1(X1,X2))!=k3_xboole_0(k1_relat_1(X3),X4)|~v1_funct_1(k7_relat_1(X1,X2))|~v1_funct_1(X3)|~v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_18, c_0_25])).
cnf(c_0_34, negated_conjecture, (k3_xboole_0(k1_relat_1(esk2_0),X1)=k1_relat_1(k7_relat_1(esk2_0,X1))), inference(rw,[status(thm)],[c_0_13, c_0_19])).
cnf(c_0_35, negated_conjecture, (k7_relat_1(esk1_0,X1)=k7_relat_1(esk2_0,esk3_0)|k1_relat_1(k7_relat_1(esk2_0,esk3_0))!=k1_relat_1(k7_relat_1(esk2_0,X1))|~v1_funct_1(k7_relat_1(esk2_0,esk3_0))|~v1_relat_1(k7_relat_1(esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_10]), c_0_34]), c_0_31]), c_0_14]), c_0_11])])).
cnf(c_0_36, negated_conjecture, (k7_relat_1(esk1_0,X1)=k7_relat_1(esk2_0,esk3_0)|k1_relat_1(k7_relat_1(esk2_0,esk3_0))!=k1_relat_1(k7_relat_1(esk2_0,X1))|~v1_relat_1(k7_relat_1(esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_23]), c_0_28]), c_0_14])])).
cnf(c_0_37, negated_conjecture, (k7_relat_1(esk1_0,esk3_0)!=k7_relat_1(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_38, negated_conjecture, (~v1_relat_1(k7_relat_1(esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_36]), c_0_37])).
cnf(c_0_39, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_22]), c_0_28]), c_0_14])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 40
# Proof object clause steps            : 27
# Proof object formula steps           : 13
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 31
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 16
# Processed clauses                    : 82
# ...of these trivial                  : 0
# ...subsumed                          : 13
# ...remaining for further processing  : 69
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 1
# Generated clauses                    : 99
# ...of the previous two non-trivial   : 90
# Contextual simplify-reflections      : 4
# Paramodulations                      : 93
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 50
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 41
# Current number of unprocessed clauses: 40
# ...number of literals in the above   : 359
# Current number of archived formulas  : 0
# Current number of archived clauses   : 19
# Clause-clause subsumption calls (NU) : 536
# Rec. Clause-clause subsumption calls : 131
# Non-unit clause-clause subsumptions  : 19
# Unit Clause-clause subsumption calls : 3
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5112

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.006 s
# Total time               : 0.012 s
# Maximum resident set size: 28740 pages
