# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/01cbab08650145459aa4084806c3b39e.E_input', d5_xboole_0)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/01cbab08650145459aa4084806c3b39e.E_input', redefinition_k6_subset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/01cbab08650145459aa4084806c3b39e.E_input', t3_subset)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/01cbab08650145459aa4084806c3b39e.E_input', t4_subset_1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/01cbab08650145459aa4084806c3b39e.E_input', t7_ordinal1)).
fof(t29_yellow_6, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>r1_waybel_0(X1,X2,u1_struct_0(X1)))), file('data/MPTP/proofs/01cbab08650145459aa4084806c3b39e.E_input', t29_yellow_6)).
fof(t10_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r2_waybel_0(X1,X2,X3)<=>~(r1_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3)))))), file('data/MPTP/proofs/01cbab08650145459aa4084806c3b39e.E_input', t10_waybel_0)).
fof(d12_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r2_waybel_0(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X4,X5)&r2_hidden(k2_waybel_0(X1,X2,X5),X3))))))), file('data/MPTP/proofs/01cbab08650145459aa4084806c3b39e.E_input', d12_waybel_0)).
fof(existence_m1_subset_1, axiom, ![X1]:?[X2]:m1_subset_1(X2,X1), file('data/MPTP/proofs/01cbab08650145459aa4084806c3b39e.E_input', existence_m1_subset_1)).
fof(c_0_9, plain, ![X26, X27, X28, X29, X30, X31, X32, X33]:((((r2_hidden(X29,X26)|~r2_hidden(X29,X28)|X28!=k4_xboole_0(X26,X27))&(~r2_hidden(X29,X27)|~r2_hidden(X29,X28)|X28!=k4_xboole_0(X26,X27)))&(~r2_hidden(X30,X26)|r2_hidden(X30,X27)|r2_hidden(X30,X28)|X28!=k4_xboole_0(X26,X27)))&((~r2_hidden(esk5_3(X31,X32,X33),X33)|(~r2_hidden(esk5_3(X31,X32,X33),X31)|r2_hidden(esk5_3(X31,X32,X33),X32))|X33=k4_xboole_0(X31,X32))&((r2_hidden(esk5_3(X31,X32,X33),X31)|r2_hidden(esk5_3(X31,X32,X33),X33)|X33=k4_xboole_0(X31,X32))&(~r2_hidden(esk5_3(X31,X32,X33),X32)|r2_hidden(esk5_3(X31,X32,X33),X33)|X33=k4_xboole_0(X31,X32))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_10, plain, ![X12, X13]:k6_subset_1(X12,X13)=k4_xboole_0(X12,X13), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
fof(c_0_11, plain, ![X14, X15]:((~m1_subset_1(X14,k1_zfmisc_1(X15))|r1_tarski(X14,X15))&(~r1_tarski(X14,X15)|m1_subset_1(X14,k1_zfmisc_1(X15)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_12, plain, ![X8]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X8)), inference(variable_rename,[status(thm)],[t4_subset_1])).
cnf(c_0_13, plain, (r2_hidden(esk5_3(X1,X2,X3),X1)|r2_hidden(esk5_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, plain, ![X24, X25]:(~r2_hidden(X24,X25)|~r1_tarski(X25,X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_16, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, plain, (r2_hidden(esk5_3(X1,X2,X3),X2)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk5_3(X1,X2,X3),X3)|~r2_hidden(esk5_3(X1,X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_19, plain, (X3=k6_subset_1(X1,X2)|r2_hidden(esk5_3(X1,X2,X3),X3)|r2_hidden(esk5_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_20, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_22, plain, (X3=k6_subset_1(X1,X2)|r2_hidden(esk5_3(X1,X2,X3),X2)|~r2_hidden(esk5_3(X1,X2,X3),X3)|~r2_hidden(esk5_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_18, c_0_14])).
cnf(c_0_23, plain, (k6_subset_1(X1,X2)=X1|r2_hidden(esk5_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_19])).
fof(c_0_24, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>r1_waybel_0(X1,X2,u1_struct_0(X1))))), inference(assume_negation,[status(cth)],[t29_yellow_6])).
fof(c_0_25, plain, ![X9, X10, X11]:((~r2_waybel_0(X9,X10,X11)|~r1_waybel_0(X9,X10,k6_subset_1(u1_struct_0(X9),X11))|(v2_struct_0(X10)|~l1_waybel_0(X10,X9))|(v2_struct_0(X9)|~l1_struct_0(X9)))&(r1_waybel_0(X9,X10,k6_subset_1(u1_struct_0(X9),X11))|r2_waybel_0(X9,X10,X11)|(v2_struct_0(X10)|~l1_waybel_0(X10,X9))|(v2_struct_0(X9)|~l1_struct_0(X9)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t10_waybel_0])])])])])).
cnf(c_0_26, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27, plain, (k6_subset_1(X1,X2)=X1|r2_hidden(esk5_3(X1,X2,X1),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_23])).
fof(c_0_28, plain, ![X16, X17, X18, X19, X21, X23]:(((m1_subset_1(esk3_4(X16,X17,X18,X19),u1_struct_0(X17))|~m1_subset_1(X19,u1_struct_0(X17))|~r2_waybel_0(X16,X17,X18)|(v2_struct_0(X17)|~l1_waybel_0(X17,X16))|(v2_struct_0(X16)|~l1_struct_0(X16)))&((r1_orders_2(X17,X19,esk3_4(X16,X17,X18,X19))|~m1_subset_1(X19,u1_struct_0(X17))|~r2_waybel_0(X16,X17,X18)|(v2_struct_0(X17)|~l1_waybel_0(X17,X16))|(v2_struct_0(X16)|~l1_struct_0(X16)))&(r2_hidden(k2_waybel_0(X16,X17,esk3_4(X16,X17,X18,X19)),X18)|~m1_subset_1(X19,u1_struct_0(X17))|~r2_waybel_0(X16,X17,X18)|(v2_struct_0(X17)|~l1_waybel_0(X17,X16))|(v2_struct_0(X16)|~l1_struct_0(X16)))))&((m1_subset_1(esk4_3(X16,X17,X21),u1_struct_0(X17))|r2_waybel_0(X16,X17,X21)|(v2_struct_0(X17)|~l1_waybel_0(X17,X16))|(v2_struct_0(X16)|~l1_struct_0(X16)))&(~m1_subset_1(X23,u1_struct_0(X17))|(~r1_orders_2(X17,esk4_3(X16,X17,X21),X23)|~r2_hidden(k2_waybel_0(X16,X17,X23),X21))|r2_waybel_0(X16,X17,X21)|(v2_struct_0(X17)|~l1_waybel_0(X17,X16))|(v2_struct_0(X16)|~l1_struct_0(X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d12_waybel_0])])])])])])])).
fof(c_0_29, negated_conjecture, ((~v2_struct_0(esk1_0)&l1_struct_0(esk1_0))&((~v2_struct_0(esk2_0)&(v4_orders_2(esk2_0)&(v7_waybel_0(esk2_0)&l1_waybel_0(esk2_0,esk1_0))))&~r1_waybel_0(esk1_0,esk2_0,u1_struct_0(esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_24])])])])).
cnf(c_0_30, plain, (r1_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3))|r2_waybel_0(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, plain, (k6_subset_1(X1,k1_xboole_0)=X1), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32, plain, (r2_hidden(k2_waybel_0(X1,X2,esk3_4(X1,X2,X3,X4)),X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~r2_waybel_0(X1,X2,X3)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33, negated_conjecture, (~r1_waybel_0(esk1_0,esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34, plain, (r2_waybel_0(X1,X2,k1_xboole_0)|r1_waybel_0(X1,X2,u1_struct_0(X1))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_35, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, negated_conjecture, (l1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X1,X2,k1_xboole_0)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_26, c_0_32])).
cnf(c_0_40, negated_conjecture, (r2_waybel_0(esk1_0,esk2_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35]), c_0_36])]), c_0_37]), c_0_38])).
fof(c_0_41, plain, ![X35]:m1_subset_1(esk6_1(X35),X35), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[existence_m1_subset_1])])).
cnf(c_0_42, negated_conjecture, (~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_35]), c_0_36])]), c_0_37]), c_0_38])).
cnf(c_0_43, plain, (m1_subset_1(esk6_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_44, negated_conjecture, ($false), inference(spm,[status(thm)],[c_0_42, c_0_43]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 45
# Proof object clause steps            : 26
# Proof object formula steps           : 19
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 13
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 25
# Processed clauses                    : 118
# ...of these trivial                  : 0
# ...subsumed                          : 31
# ...remaining for further processing  : 87
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 184
# ...of the previous two non-trivial   : 120
# Contextual simplify-reflections      : 2
# Paramodulations                      : 179
# Factorizations                       : 2
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 58
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 38
# Current number of unprocessed clauses: 51
# ...number of literals in the above   : 160
# Current number of archived formulas  : 0
# Current number of archived clauses   : 27
# Clause-clause subsumption calls (NU) : 624
# Rec. Clause-clause subsumption calls : 260
# Non-unit clause-clause subsumptions  : 21
# Unit Clause-clause subsumption calls : 58
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4295

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.003 s
# Total time               : 0.012 s
# Maximum resident set size: 28760 pages
