# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t71_mcart_1, conjecture, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X8))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k7_mcart_1(X1,X2,X3,X5)))))), file('data/MPTP/proofs/020e414649f5479aa7d7c3d4e1d57dff.E_input', t71_mcart_1)).
fof(t10_mcart_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(r2_hidden(k1_mcart_1(X1),X2)&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/020e414649f5479aa7d7c3d4e1d57dff.E_input', t10_mcart_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/020e414649f5479aa7d7c3d4e1d57dff.E_input', d1_xboole_0)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/020e414649f5479aa7d7c3d4e1d57dff.E_input', d3_zfmisc_1)).
fof(t34_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4, X5, X6]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k4_mcart_1(X3,X4,X5,X6))))))), file('data/MPTP/proofs/020e414649f5479aa7d7c3d4e1d57dff.E_input', t34_mcart_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/020e414649f5479aa7d7c3d4e1d57dff.E_input', d2_subset_1)).
fof(d4_zfmisc_1, axiom, ![X1, X2, X3, X4]:k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4), file('data/MPTP/proofs/020e414649f5479aa7d7c3d4e1d57dff.E_input', d4_zfmisc_1)).
fof(t55_mcart_1, axiom, ![X1, X2, X3, X4]:((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(X4=k1_xboole_0))))<=>~(k4_zfmisc_1(X1,X2,X3,X4)=k1_xboole_0)), file('data/MPTP/proofs/020e414649f5479aa7d7c3d4e1d57dff.E_input', t55_mcart_1)).
fof(l139_zfmisc_1, axiom, ![X1, X2, X3]:~((r2_hidden(X1,k2_zfmisc_1(X2,X3))&![X4, X5]:~(k4_tarski(X4,X5)=X1))), file('data/MPTP/proofs/020e414649f5479aa7d7c3d4e1d57dff.E_input', l139_zfmisc_1)).
fof(fc1_zfmisc_1, axiom, ![X1, X2]:~(v1_xboole_0(k4_tarski(X1,X2))), file('data/MPTP/proofs/020e414649f5479aa7d7c3d4e1d57dff.E_input', fc1_zfmisc_1)).
fof(d1_relat_1, axiom, ![X1]:(v1_relat_1(X1)<=>![X2]:~((r2_hidden(X2,X1)&![X3, X4]:~(X2=k4_tarski(X3,X4))))), file('data/MPTP/proofs/020e414649f5479aa7d7c3d4e1d57dff.E_input', d1_relat_1)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/020e414649f5479aa7d7c3d4e1d57dff.E_input', t7_mcart_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/020e414649f5479aa7d7c3d4e1d57dff.E_input', fc1_xboole_0)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/020e414649f5479aa7d7c3d4e1d57dff.E_input', fc6_relat_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/020e414649f5479aa7d7c3d4e1d57dff.E_input', d3_mcart_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/020e414649f5479aa7d7c3d4e1d57dff.E_input', t1_subset)).
fof(t50_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))&(k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))&k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4))))))))), file('data/MPTP/proofs/020e414649f5479aa7d7c3d4e1d57dff.E_input', t50_mcart_1)).
fof(c_0_17, negated_conjecture, ~(![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X8))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k7_mcart_1(X1,X2,X3,X5))))))), inference(assume_negation,[status(cth)],[t71_mcart_1])).
fof(c_0_18, plain, ![X17, X18, X19]:((r2_hidden(k1_mcart_1(X17),X18)|~r2_hidden(X17,k2_zfmisc_1(X18,X19)))&(r2_hidden(k2_mcart_1(X17),X19)|~r2_hidden(X17,k2_zfmisc_1(X18,X19)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_mcart_1])])])).
fof(c_0_19, plain, ![X20, X21, X22]:((~v1_xboole_0(X20)|~r2_hidden(X21,X20))&(r2_hidden(esk6_1(X22),X22)|v1_xboole_0(X22))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_20, negated_conjecture, ![X14, X15, X16]:(m1_subset_1(esk5_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))&((~m1_subset_1(X14,esk1_0)|(~m1_subset_1(X15,esk2_0)|(~m1_subset_1(X16,esk3_0)|(esk5_0!=k3_mcart_1(X14,X15,X16)|esk4_0=X16))))&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&esk4_0!=k7_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])])).
fof(c_0_21, plain, ![X24, X25, X26]:k3_zfmisc_1(X24,X25,X26)=k2_zfmisc_1(k2_zfmisc_1(X24,X25),X26), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_22, plain, ![X27, X29, X30, X31, X32]:((r2_hidden(esk7_1(X27),X27)|X27=k1_xboole_0)&((~r2_hidden(X29,X27)|esk7_1(X27)!=k4_mcart_1(X29,X30,X31,X32)|X27=k1_xboole_0)&(~r2_hidden(X30,X27)|esk7_1(X27)!=k4_mcart_1(X29,X30,X31,X32)|X27=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_mcart_1])])])])])])])).
cnf(c_0_23, plain, (r2_hidden(k1_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (r2_hidden(esk6_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_25, plain, ![X33, X34]:(((~m1_subset_1(X34,X33)|r2_hidden(X34,X33)|v1_xboole_0(X33))&(~r2_hidden(X34,X33)|m1_subset_1(X34,X33)|v1_xboole_0(X33)))&((~m1_subset_1(X34,X33)|v1_xboole_0(X34)|~v1_xboole_0(X33))&(~v1_xboole_0(X34)|m1_subset_1(X34,X33)|~v1_xboole_0(X33)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
cnf(c_0_26, negated_conjecture, (m1_subset_1(esk5_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_28, plain, ![X37, X38, X39, X40]:k4_zfmisc_1(X37,X38,X39,X40)=k2_zfmisc_1(k3_zfmisc_1(X37,X38,X39),X40), inference(variable_rename,[status(thm)],[d4_zfmisc_1])).
cnf(c_0_29, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_30, plain, (r2_hidden(esk7_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|r2_hidden(k1_mcart_1(esk6_1(k2_zfmisc_1(X1,X2))),X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_32, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, negated_conjecture, (m1_subset_1(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(rw,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_34, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_35, plain, ![X41, X42, X43, X44]:((X41=k1_xboole_0|(X42=k1_xboole_0|(X43=k1_xboole_0|X44=k1_xboole_0))|k4_zfmisc_1(X41,X42,X43,X44)!=k1_xboole_0)&((X41!=k1_xboole_0|k4_zfmisc_1(X41,X42,X43,X44)=k1_xboole_0)&((X42!=k1_xboole_0|k4_zfmisc_1(X41,X42,X43,X44)=k1_xboole_0)&((X43!=k1_xboole_0|k4_zfmisc_1(X41,X42,X43,X44)=k1_xboole_0)&(X44!=k1_xboole_0|k4_zfmisc_1(X41,X42,X43,X44)=k1_xboole_0))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t55_mcart_1])])])])).
cnf(c_0_36, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_38, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_29, c_0_31])).
cnf(c_0_39, negated_conjecture, (v1_xboole_0(esk5_0)|~v1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_40, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_34, c_0_33])).
cnf(c_0_41, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|k4_zfmisc_1(X1,X2,X3,X4)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_42, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_36, c_0_27])).
cnf(c_0_43, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
fof(c_0_44, plain, ![X45, X46, X47]:(~r2_hidden(X45,k2_zfmisc_1(X46,X47))|k4_tarski(esk8_1(X45),esk9_1(X45))=X45), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l139_zfmisc_1])])])])])])).
cnf(c_0_45, negated_conjecture, (v1_xboole_0(esk5_0)|r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_39, c_0_40])).
fof(c_0_46, plain, ![X50, X51]:~v1_xboole_0(k4_tarski(X50,X51)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_zfmisc_1])])).
fof(c_0_47, plain, ![X52, X53, X56, X58, X59]:((~v1_relat_1(X52)|(~r2_hidden(X53,X52)|X53=k4_tarski(esk10_2(X52,X53),esk11_2(X52,X53))))&((r2_hidden(esk12_1(X56),X56)|v1_relat_1(X56))&(esk12_1(X56)!=k4_tarski(X58,X59)|v1_relat_1(X56)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_relat_1])])])])])])])).
cnf(c_0_48, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_49, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),X1)=k1_xboole_0|r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_43, c_0_40])).
cnf(c_0_50, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_51, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_52, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_53, plain, ![X60, X61]:(k1_mcart_1(k4_tarski(X60,X61))=X60&k2_mcart_1(k4_tarski(X60,X61))=X61), inference(variable_rename,[status(thm)],[t7_mcart_1])).
cnf(c_0_54, plain, (k4_tarski(esk8_1(X1),esk9_1(X1))=X1|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_55, negated_conjecture, (v1_xboole_0(esk5_0)|r2_hidden(k1_mcart_1(esk5_0),k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_23, c_0_45])).
cnf(c_0_56, plain, (~v1_xboole_0(k4_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_57, plain, (X2=k4_tarski(esk10_2(X1,X2),esk11_2(X1,X2))|~v1_relat_1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_58, negated_conjecture, (X1=k1_xboole_0|r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50]), c_0_51]), c_0_52])).
cnf(c_0_59, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
fof(c_0_60, plain, ![X62, X63]:v1_relat_1(k2_zfmisc_1(X62,X63)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_61, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_62, negated_conjecture, (k4_tarski(esk8_1(k1_mcart_1(esk5_0)),esk9_1(k1_mcart_1(esk5_0)))=k1_mcart_1(esk5_0)|v1_xboole_0(esk5_0)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_63, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_64, plain, (~v1_relat_1(X1)|~v1_xboole_0(X2)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_65, negated_conjecture, (r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_58]), c_0_59])])).
cnf(c_0_66, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_60])).
fof(c_0_67, plain, ![X64, X65, X66]:k3_mcart_1(X64,X65,X66)=k4_tarski(k4_tarski(X64,X65),X66), inference(variable_rename,[status(thm)],[d3_mcart_1])).
cnf(c_0_68, negated_conjecture, (esk8_1(k1_mcart_1(esk5_0))=k1_mcart_1(k1_mcart_1(esk5_0))|v1_xboole_0(esk5_0)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_69, negated_conjecture, (esk9_1(k1_mcart_1(esk5_0))=k2_mcart_1(k1_mcart_1(esk5_0))|v1_xboole_0(esk5_0)), inference(spm,[status(thm)],[c_0_63, c_0_62])).
cnf(c_0_70, negated_conjecture, (~v1_xboole_0(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_66])])).
cnf(c_0_71, negated_conjecture, (esk4_0=X3|~m1_subset_1(X1,esk1_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X3,esk3_0)|esk5_0!=k3_mcart_1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_72, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_73, negated_conjecture, (k4_tarski(k1_mcart_1(k1_mcart_1(esk5_0)),esk9_1(k1_mcart_1(esk5_0)))=k1_mcart_1(esk5_0)|v1_xboole_0(esk5_0)), inference(spm,[status(thm)],[c_0_62, c_0_68])).
cnf(c_0_74, negated_conjecture, (esk9_1(k1_mcart_1(esk5_0))=k2_mcart_1(k1_mcart_1(esk5_0))), inference(sr,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_75, plain, (r2_hidden(k2_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_76, negated_conjecture, (esk4_0=X3|esk5_0!=k4_tarski(k4_tarski(X1,X2),X3)|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk1_0)), inference(rw,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_77, negated_conjecture, (k4_tarski(k1_mcart_1(k1_mcart_1(esk5_0)),k2_mcart_1(k1_mcart_1(esk5_0)))=k1_mcart_1(esk5_0)), inference(rw,[status(thm)],[inference(sr,[status(thm)],[c_0_73, c_0_70]), c_0_74])).
fof(c_0_78, plain, ![X67, X68]:(~r2_hidden(X67,X68)|m1_subset_1(X67,X68)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_79, negated_conjecture, (v1_xboole_0(esk5_0)|r2_hidden(k2_mcart_1(k1_mcart_1(esk5_0)),esk2_0)), inference(spm,[status(thm)],[c_0_75, c_0_55])).
cnf(c_0_80, negated_conjecture, (k4_tarski(esk8_1(esk5_0),esk9_1(esk5_0))=esk5_0|v1_xboole_0(esk5_0)), inference(spm,[status(thm)],[c_0_54, c_0_45])).
fof(c_0_81, plain, ![X69, X70, X71, X72]:((k5_mcart_1(X69,X70,X71,X72)=k1_mcart_1(k1_mcart_1(X72))|~m1_subset_1(X72,k3_zfmisc_1(X69,X70,X71))|X71=k1_xboole_0|X70=k1_xboole_0|X69=k1_xboole_0)&((k6_mcart_1(X69,X70,X71,X72)=k2_mcart_1(k1_mcart_1(X72))|~m1_subset_1(X72,k3_zfmisc_1(X69,X70,X71))|X71=k1_xboole_0|X70=k1_xboole_0|X69=k1_xboole_0)&(k7_mcart_1(X69,X70,X71,X72)=k2_mcart_1(X72)|~m1_subset_1(X72,k3_zfmisc_1(X69,X70,X71))|X71=k1_xboole_0|X70=k1_xboole_0|X69=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_mcart_1])])])])])).
cnf(c_0_82, negated_conjecture, (esk4_0=X1|k4_tarski(k1_mcart_1(esk5_0),X1)!=esk5_0|~m1_subset_1(k2_mcart_1(k1_mcart_1(esk5_0)),esk2_0)|~m1_subset_1(k1_mcart_1(k1_mcart_1(esk5_0)),esk1_0)|~m1_subset_1(X1,esk3_0)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_83, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_84, negated_conjecture, (r2_hidden(k2_mcart_1(k1_mcart_1(esk5_0)),esk2_0)), inference(sr,[status(thm)],[c_0_79, c_0_70])).
cnf(c_0_85, negated_conjecture, (v1_xboole_0(esk5_0)|r2_hidden(k1_mcart_1(k1_mcart_1(esk5_0)),esk1_0)), inference(spm,[status(thm)],[c_0_23, c_0_55])).
cnf(c_0_86, negated_conjecture, (esk8_1(esk5_0)=k1_mcart_1(esk5_0)|v1_xboole_0(esk5_0)), inference(spm,[status(thm)],[c_0_61, c_0_80])).
cnf(c_0_87, negated_conjecture, (esk9_1(esk5_0)=k2_mcart_1(esk5_0)|v1_xboole_0(esk5_0)), inference(spm,[status(thm)],[c_0_63, c_0_80])).
cnf(c_0_88, plain, (k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_89, negated_conjecture, (esk4_0=X1|k4_tarski(k1_mcart_1(esk5_0),X1)!=esk5_0|~m1_subset_1(k1_mcart_1(k1_mcart_1(esk5_0)),esk1_0)|~m1_subset_1(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_84])])).
cnf(c_0_90, negated_conjecture, (r2_hidden(k1_mcart_1(k1_mcart_1(esk5_0)),esk1_0)), inference(sr,[status(thm)],[c_0_85, c_0_70])).
cnf(c_0_91, negated_conjecture, (esk8_1(esk5_0)=k1_mcart_1(esk5_0)), inference(sr,[status(thm)],[c_0_86, c_0_70])).
cnf(c_0_92, negated_conjecture, (esk9_1(esk5_0)=k2_mcart_1(esk5_0)), inference(sr,[status(thm)],[c_0_87, c_0_70])).
cnf(c_0_93, negated_conjecture, (esk4_0!=k7_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_94, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_88, c_0_27])).
cnf(c_0_95, negated_conjecture, (esk4_0=X1|k4_tarski(k1_mcart_1(esk5_0),X1)!=esk5_0|~m1_subset_1(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_83]), c_0_90])])).
cnf(c_0_96, negated_conjecture, (k4_tarski(k1_mcart_1(esk5_0),k2_mcart_1(esk5_0))=esk5_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_65]), c_0_91]), c_0_92])).
cnf(c_0_97, negated_conjecture, (k2_mcart_1(esk5_0)!=esk4_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_33])]), c_0_50]), c_0_51]), c_0_52])).
cnf(c_0_98, negated_conjecture, (X1=k1_xboole_0|r2_hidden(k2_mcart_1(esk5_0),esk3_0)), inference(spm,[status(thm)],[c_0_75, c_0_58])).
cnf(c_0_99, negated_conjecture, (~m1_subset_1(k2_mcart_1(esk5_0),esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_97])).
cnf(c_0_100, negated_conjecture, (r2_hidden(k2_mcart_1(esk5_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_98]), c_0_59])])).
cnf(c_0_101, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_83]), c_0_100])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 102
# Proof object clause steps            : 68
# Proof object formula steps           : 34
# Proof object conjectures             : 43
# Proof object clause conjectures      : 40
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 31
# Proof object simplifying inferences  : 35
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 36
# Processed clauses                    : 508
# ...of these trivial                  : 18
# ...subsumed                          : 239
# ...remaining for further processing  : 251
# Other redundant clauses eliminated   : 19
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 37
# Backward-rewritten                   : 18
# Generated clauses                    : 2223
# ...of the previous two non-trivial   : 1145
# Contextual simplify-reflections      : 1
# Paramodulations                      : 2147
# Factorizations                       : 42
# Equation resolutions                 : 19
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 143
#    Positive orientable unit clauses  : 20
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 15
#    Non-unit-clauses                  : 108
# Current number of unprocessed clauses: 283
# ...number of literals in the above   : 1164
# Current number of archived formulas  : 0
# Current number of archived clauses   : 107
# Clause-clause subsumption calls (NU) : 1847
# Rec. Clause-clause subsumption calls : 1358
# Non-unit clause-clause subsumptions  : 225
# Unit Clause-clause subsumption calls : 230
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 23
# BW rewrite match successes           : 16
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 23766

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.006 s
# Total time               : 0.023 s
# Maximum resident set size: 28768 pages
