# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C47_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t17_yellow_1, conjecture, ![X1, X2]:(m1_subset_1(X2,u1_struct_0(k3_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k3_yellow_1(X1)))=>(k13_lattice3(k3_yellow_1(X1),X2,X3)=k2_xboole_0(X2,X3)&k12_lattice3(k3_yellow_1(X1),X2,X3)=k3_xboole_0(X2,X3)))), file('data/MPTP/proofs/0278b9e68a454a9e81506a48bb0f0c1f.E_input', t17_yellow_1)).
fof(dt_k2_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&l1_orders_2(k2_yellow_1(X1))), file('data/MPTP/proofs/0278b9e68a454a9e81506a48bb0f0c1f.E_input', dt_k2_yellow_1)).
fof(d1_yellow_1, axiom, ![X1]:k2_yellow_1(X1)=g1_orders_2(X1,k1_yellow_1(X1)), file('data/MPTP/proofs/0278b9e68a454a9e81506a48bb0f0c1f.E_input', d1_yellow_1)).
fof(t4_yellow_1, axiom, ![X1]:k3_yellow_1(X1)=k2_yellow_1(k9_setfam_1(X1)), file('data/MPTP/proofs/0278b9e68a454a9e81506a48bb0f0c1f.E_input', t4_yellow_1)).
fof(d2_yellow_1, axiom, ![X1]:k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1)), file('data/MPTP/proofs/0278b9e68a454a9e81506a48bb0f0c1f.E_input', d2_yellow_1)).
fof(fc5_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&(v3_orders_2(k2_yellow_1(X1))&(v4_orders_2(k2_yellow_1(X1))&v5_orders_2(k2_yellow_1(X1))))), file('data/MPTP/proofs/0278b9e68a454a9e81506a48bb0f0c1f.E_input', fc5_yellow_1)).
fof(t8_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r2_hidden(k2_xboole_0(X2,X3),X1)=>k10_lattice3(k2_yellow_1(X1),X2,X3)=k2_xboole_0(X2,X3))))), file('data/MPTP/proofs/0278b9e68a454a9e81506a48bb0f0c1f.E_input', t8_yellow_1)).
fof(redefinition_k13_lattice3, axiom, ![X1, X2, X3]:(((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k13_lattice3(X1,X2,X3)=k10_lattice3(X1,X2,X3)), file('data/MPTP/proofs/0278b9e68a454a9e81506a48bb0f0c1f.E_input', redefinition_k13_lattice3)).
fof(t7_boole, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2))), file('data/MPTP/proofs/0278b9e68a454a9e81506a48bb0f0c1f.E_input', t7_boole)).
fof(l19_yellow_1, axiom, ![X1, X2]:(m1_subset_1(X2,u1_struct_0(k3_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k3_yellow_1(X1)))=>(r2_hidden(k3_xboole_0(X2,X3),k9_setfam_1(X1))&r2_hidden(k2_xboole_0(X2,X3),k9_setfam_1(X1))))), file('data/MPTP/proofs/0278b9e68a454a9e81506a48bb0f0c1f.E_input', l19_yellow_1)).
fof(t9_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r2_hidden(k3_xboole_0(X2,X3),X1)=>k11_lattice3(k2_yellow_1(X1),X2,X3)=k3_xboole_0(X2,X3))))), file('data/MPTP/proofs/0278b9e68a454a9e81506a48bb0f0c1f.E_input', t9_yellow_1)).
fof(redefinition_k12_lattice3, axiom, ![X1, X2, X3]:(((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)), file('data/MPTP/proofs/0278b9e68a454a9e81506a48bb0f0c1f.E_input', redefinition_k12_lattice3)).
fof(fc1_yellow_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>(v1_orders_2(k3_lattice3(X1))&(v3_orders_2(k3_lattice3(X1))&(v4_orders_2(k3_lattice3(X1))&(v5_orders_2(k3_lattice3(X1))&(v1_lattice3(k3_lattice3(X1))&v2_lattice3(k3_lattice3(X1)))))))), file('data/MPTP/proofs/0278b9e68a454a9e81506a48bb0f0c1f.E_input', fc1_yellow_1)).
fof(dt_k1_lattice3, axiom, ![X1]:(v3_lattices(k1_lattice3(X1))&l3_lattices(k1_lattice3(X1))), file('data/MPTP/proofs/0278b9e68a454a9e81506a48bb0f0c1f.E_input', dt_k1_lattice3)).
fof(fc2_lattice3, axiom, ![X1]:(v3_lattices(k1_lattice3(X1))&v10_lattices(k1_lattice3(X1))), file('data/MPTP/proofs/0278b9e68a454a9e81506a48bb0f0c1f.E_input', fc2_lattice3)).
fof(fc1_lattice3, axiom, ![X1]:(~(v2_struct_0(k1_lattice3(X1)))&v3_lattices(k1_lattice3(X1))), file('data/MPTP/proofs/0278b9e68a454a9e81506a48bb0f0c1f.E_input', fc1_lattice3)).
fof(c_0_16, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,u1_struct_0(k3_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k3_yellow_1(X1)))=>(k13_lattice3(k3_yellow_1(X1),X2,X3)=k2_xboole_0(X2,X3)&k12_lattice3(k3_yellow_1(X1),X2,X3)=k3_xboole_0(X2,X3))))), inference(assume_negation,[status(cth)],[t17_yellow_1])).
fof(c_0_17, plain, ![X26]:(v1_orders_2(k2_yellow_1(X26))&l1_orders_2(k2_yellow_1(X26))), inference(variable_rename,[status(thm)],[dt_k2_yellow_1])).
fof(c_0_18, plain, ![X15]:k2_yellow_1(X15)=g1_orders_2(X15,k1_yellow_1(X15)), inference(variable_rename,[status(thm)],[d1_yellow_1])).
fof(c_0_19, plain, ![X27]:k3_yellow_1(X27)=k2_yellow_1(k9_setfam_1(X27)), inference(variable_rename,[status(thm)],[t4_yellow_1])).
fof(c_0_20, plain, ![X31]:k3_yellow_1(X31)=k3_lattice3(k1_lattice3(X31)), inference(variable_rename,[status(thm)],[d2_yellow_1])).
fof(c_0_21, plain, ![X22]:(v1_orders_2(k2_yellow_1(X22))&(v3_orders_2(k2_yellow_1(X22))&(v4_orders_2(k2_yellow_1(X22))&v5_orders_2(k2_yellow_1(X22))))), inference(variable_rename,[status(thm)],[fc5_yellow_1])).
fof(c_0_22, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(k3_yellow_1(esk1_0)))&(k13_lattice3(k3_yellow_1(esk1_0),esk2_0,esk3_0)!=k2_xboole_0(esk2_0,esk3_0)|k12_lattice3(k3_yellow_1(esk1_0),esk2_0,esk3_0)!=k3_xboole_0(esk2_0,esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).
cnf(c_0_23, plain, (l1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (k2_yellow_1(X1)=g1_orders_2(X1,k1_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (k3_yellow_1(X1)=k2_yellow_1(k9_setfam_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (v5_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_28, plain, ![X12, X13, X14]:(v1_xboole_0(X12)|(~m1_subset_1(X13,u1_struct_0(k2_yellow_1(X12)))|(~m1_subset_1(X14,u1_struct_0(k2_yellow_1(X12)))|(~r2_hidden(k2_xboole_0(X13,X14),X12)|k10_lattice3(k2_yellow_1(X12),X13,X14)=k2_xboole_0(X13,X14))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_yellow_1])])])])).
cnf(c_0_29, negated_conjecture, (k13_lattice3(k3_yellow_1(esk1_0),esk2_0,esk3_0)!=k2_xboole_0(esk2_0,esk3_0)|k12_lattice3(k3_yellow_1(esk1_0),esk2_0,esk3_0)!=k3_xboole_0(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_30, plain, ![X28, X29, X30]:(~v5_orders_2(X28)|(~v1_lattice3(X28)|~l1_orders_2(X28))|(~m1_subset_1(X29,u1_struct_0(X28))|~m1_subset_1(X30,u1_struct_0(X28)))|k13_lattice3(X28,X29,X30)=k10_lattice3(X28,X29,X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k13_lattice3])])).
cnf(c_0_31, plain, (l1_orders_2(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_32, plain, (k3_lattice3(k1_lattice3(X1))=g1_orders_2(k9_setfam_1(X1),k1_yellow_1(k9_setfam_1(X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_26]), c_0_24])).
cnf(c_0_33, plain, (v5_orders_2(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_27, c_0_24])).
cnf(c_0_34, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k3_yellow_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_35, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_36, plain, (v1_xboole_0(X1)|k10_lattice3(k2_yellow_1(X1),X2,X3)=k2_xboole_0(X2,X3)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))|~r2_hidden(k2_xboole_0(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_37, plain, ![X24, X25]:(~r2_hidden(X24,X25)|~v1_xboole_0(X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
fof(c_0_38, plain, ![X16, X17, X18]:((r2_hidden(k3_xboole_0(X17,X18),k9_setfam_1(X16))|~m1_subset_1(X18,u1_struct_0(k3_yellow_1(X16)))|~m1_subset_1(X17,u1_struct_0(k3_yellow_1(X16))))&(r2_hidden(k2_xboole_0(X17,X18),k9_setfam_1(X16))|~m1_subset_1(X18,u1_struct_0(k3_yellow_1(X16)))|~m1_subset_1(X17,u1_struct_0(k3_yellow_1(X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l19_yellow_1])])])])).
fof(c_0_39, plain, ![X19, X20, X21]:(v1_xboole_0(X19)|(~m1_subset_1(X20,u1_struct_0(k2_yellow_1(X19)))|(~m1_subset_1(X21,u1_struct_0(k2_yellow_1(X19)))|(~r2_hidden(k3_xboole_0(X20,X21),X19)|k11_lattice3(k2_yellow_1(X19),X20,X21)=k3_xboole_0(X20,X21))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_yellow_1])])])])).
cnf(c_0_40, negated_conjecture, (k13_lattice3(k3_lattice3(k1_lattice3(esk1_0)),esk2_0,esk3_0)!=k2_xboole_0(esk2_0,esk3_0)|k12_lattice3(k3_lattice3(k1_lattice3(esk1_0)),esk2_0,esk3_0)!=k3_xboole_0(esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_26]), c_0_26])).
cnf(c_0_41, plain, (k13_lattice3(X1,X2,X3)=k10_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_42, plain, (l1_orders_2(k3_lattice3(k1_lattice3(X1)))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_43, plain, (v5_orders_2(k3_lattice3(k1_lattice3(X1)))), inference(spm,[status(thm)],[c_0_33, c_0_32])).
cnf(c_0_44, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k3_lattice3(k1_lattice3(esk1_0))))), inference(rw,[status(thm)],[c_0_34, c_0_26])).
cnf(c_0_45, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(k3_lattice3(k1_lattice3(esk1_0))))), inference(rw,[status(thm)],[c_0_35, c_0_26])).
fof(c_0_46, plain, ![X8, X9, X10]:(~v5_orders_2(X8)|(~v2_lattice3(X8)|~l1_orders_2(X8))|(~m1_subset_1(X9,u1_struct_0(X8))|~m1_subset_1(X10,u1_struct_0(X8)))|k12_lattice3(X8,X9,X10)=k11_lattice3(X8,X9,X10)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k12_lattice3])])).
cnf(c_0_47, plain, (k10_lattice3(g1_orders_2(X1,k1_yellow_1(X1)),X2,X3)=k2_xboole_0(X2,X3)|v1_xboole_0(X1)|~r2_hidden(k2_xboole_0(X2,X3),X1)|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_24]), c_0_24]), c_0_24])).
cnf(c_0_48, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_49, plain, (r2_hidden(k2_xboole_0(X1,X2),k9_setfam_1(X3))|~m1_subset_1(X2,u1_struct_0(k3_yellow_1(X3)))|~m1_subset_1(X1,u1_struct_0(k3_yellow_1(X3)))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, plain, (v1_xboole_0(X1)|k11_lattice3(k2_yellow_1(X1),X2,X3)=k3_xboole_0(X2,X3)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))|~r2_hidden(k3_xboole_0(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51, negated_conjecture, (k10_lattice3(k3_lattice3(k1_lattice3(esk1_0)),esk2_0,esk3_0)!=k2_xboole_0(esk2_0,esk3_0)|k12_lattice3(k3_lattice3(k1_lattice3(esk1_0)),esk2_0,esk3_0)!=k3_xboole_0(esk2_0,esk3_0)|~v1_lattice3(k3_lattice3(k1_lattice3(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42]), c_0_43]), c_0_44]), c_0_45])])).
cnf(c_0_52, plain, (k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_53, plain, (k10_lattice3(g1_orders_2(X1,k1_yellow_1(X1)),X2,X3)=k2_xboole_0(X2,X3)|~r2_hidden(k2_xboole_0(X2,X3),X1)|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))), inference(csr,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_54, plain, (r2_hidden(k2_xboole_0(X1,X2),k9_setfam_1(X3))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(X3))))|~m1_subset_1(X1,u1_struct_0(k3_lattice3(k1_lattice3(X3))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_26]), c_0_26])).
cnf(c_0_55, plain, (k11_lattice3(g1_orders_2(X1,k1_yellow_1(X1)),X2,X3)=k3_xboole_0(X2,X3)|v1_xboole_0(X1)|~r2_hidden(k3_xboole_0(X2,X3),X1)|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_24]), c_0_24]), c_0_24])).
cnf(c_0_56, plain, (r2_hidden(k3_xboole_0(X1,X2),k9_setfam_1(X3))|~m1_subset_1(X2,u1_struct_0(k3_yellow_1(X3)))|~m1_subset_1(X1,u1_struct_0(k3_yellow_1(X3)))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_57, negated_conjecture, (k10_lattice3(k3_lattice3(k1_lattice3(esk1_0)),esk2_0,esk3_0)!=k2_xboole_0(esk2_0,esk3_0)|k11_lattice3(k3_lattice3(k1_lattice3(esk1_0)),esk2_0,esk3_0)!=k3_xboole_0(esk2_0,esk3_0)|~v2_lattice3(k3_lattice3(k1_lattice3(esk1_0)))|~v1_lattice3(k3_lattice3(k1_lattice3(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_42]), c_0_43]), c_0_44]), c_0_45])])).
cnf(c_0_58, plain, (k10_lattice3(k3_lattice3(k1_lattice3(X1)),X2,X3)=k2_xboole_0(X2,X3)|~m1_subset_1(X3,u1_struct_0(k3_lattice3(k1_lattice3(X1))))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(X1))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_32]), c_0_54])).
cnf(c_0_59, plain, (k11_lattice3(g1_orders_2(X1,k1_yellow_1(X1)),X2,X3)=k3_xboole_0(X2,X3)|~r2_hidden(k3_xboole_0(X2,X3),X1)|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))), inference(csr,[status(thm)],[c_0_55, c_0_48])).
cnf(c_0_60, plain, (r2_hidden(k3_xboole_0(X1,X2),k9_setfam_1(X3))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(X3))))|~m1_subset_1(X1,u1_struct_0(k3_lattice3(k1_lattice3(X3))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_26]), c_0_26])).
cnf(c_0_61, negated_conjecture, (k11_lattice3(k3_lattice3(k1_lattice3(esk1_0)),esk2_0,esk3_0)!=k3_xboole_0(esk2_0,esk3_0)|~v2_lattice3(k3_lattice3(k1_lattice3(esk1_0)))|~v1_lattice3(k3_lattice3(k1_lattice3(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_44]), c_0_45])])).
cnf(c_0_62, plain, (k11_lattice3(k3_lattice3(k1_lattice3(X1)),X2,X3)=k3_xboole_0(X2,X3)|~m1_subset_1(X3,u1_struct_0(k3_lattice3(k1_lattice3(X1))))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(X1))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_32]), c_0_60])).
fof(c_0_63, plain, ![X7]:((v1_orders_2(k3_lattice3(X7))|(v2_struct_0(X7)|(~v10_lattices(X7)|~l3_lattices(X7))))&((v3_orders_2(k3_lattice3(X7))|(v2_struct_0(X7)|(~v10_lattices(X7)|~l3_lattices(X7))))&((v4_orders_2(k3_lattice3(X7))|(v2_struct_0(X7)|(~v10_lattices(X7)|~l3_lattices(X7))))&((v5_orders_2(k3_lattice3(X7))|(v2_struct_0(X7)|(~v10_lattices(X7)|~l3_lattices(X7))))&((v1_lattice3(k3_lattice3(X7))|(v2_struct_0(X7)|(~v10_lattices(X7)|~l3_lattices(X7))))&(v2_lattice3(k3_lattice3(X7))|(v2_struct_0(X7)|(~v10_lattices(X7)|~l3_lattices(X7))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_yellow_1])])])])).
fof(c_0_64, plain, ![X11]:(v3_lattices(k1_lattice3(X11))&l3_lattices(k1_lattice3(X11))), inference(variable_rename,[status(thm)],[dt_k1_lattice3])).
fof(c_0_65, plain, ![X32]:(v3_lattices(k1_lattice3(X32))&v10_lattices(k1_lattice3(X32))), inference(variable_rename,[status(thm)],[fc2_lattice3])).
fof(c_0_66, plain, ![X23]:(~v2_struct_0(k1_lattice3(X23))&v3_lattices(k1_lattice3(X23))), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_lattice3])])).
cnf(c_0_67, negated_conjecture, (~v2_lattice3(k3_lattice3(k1_lattice3(esk1_0)))|~v1_lattice3(k3_lattice3(k1_lattice3(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_44]), c_0_45])])).
cnf(c_0_68, plain, (v2_lattice3(k3_lattice3(X1))|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_69, plain, (l3_lattices(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_70, plain, (v10_lattices(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_71, plain, (~v2_struct_0(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_72, negated_conjecture, (~v1_lattice3(k3_lattice3(k1_lattice3(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_69]), c_0_70])]), c_0_71])).
cnf(c_0_73, plain, (v1_lattice3(k3_lattice3(X1))|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_74, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_69]), c_0_70])]), c_0_71]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 75
# Proof object clause steps            : 42
# Proof object formula steps           : 33
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 46
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 29
# Processed clauses                    : 73
# ...of these trivial                  : 3
# ...subsumed                          : 3
# ...remaining for further processing  : 67
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 0
# Generated clauses                    : 19
# ...of the previous two non-trivial   : 18
# Contextual simplify-reflections      : 4
# Paramodulations                      : 19
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 39
#    Positive orientable unit clauses  : 16
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 20
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 30
# Clause-clause subsumption calls (NU) : 120
# Rec. Clause-clause subsumption calls : 80
# Non-unit clause-clause subsumptions  : 7
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2707

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28704 pages
