# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/02a8da7181ed47f983245caac44aeda0.E_input', d4_xboole_0)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/02a8da7181ed47f983245caac44aeda0.E_input', t48_xboole_1)).
fof(t72_zfmisc_1, axiom, ![X1, X2, X3]:(k4_xboole_0(k2_tarski(X1,X2),X3)=k2_tarski(X1,X2)<=>(~(r2_hidden(X1,X3))&~(r2_hidden(X2,X3)))), file('data/MPTP/proofs/02a8da7181ed47f983245caac44aeda0.E_input', t72_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/02a8da7181ed47f983245caac44aeda0.E_input', t70_enumset1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/02a8da7181ed47f983245caac44aeda0.E_input', t3_boole)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/02a8da7181ed47f983245caac44aeda0.E_input', d1_tarski)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/02a8da7181ed47f983245caac44aeda0.E_input', t69_enumset1)).
fof(t29_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4, X5]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k3_mcart_1(X3,X4,X5))))))), file('data/MPTP/proofs/02a8da7181ed47f983245caac44aeda0.E_input', t29_mcart_1)).
fof(d1_enumset1, axiom, ![X1, X2, X3, X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3)))))), file('data/MPTP/proofs/02a8da7181ed47f983245caac44aeda0.E_input', d1_enumset1)).
fof(t51_mcart_1, conjecture, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(~(X4=k5_mcart_1(X1,X2,X3,X4))&(~(X4=k6_mcart_1(X1,X2,X3,X4))&~(X4=k7_mcart_1(X1,X2,X3,X4)))))))))), file('data/MPTP/proofs/02a8da7181ed47f983245caac44aeda0.E_input', t51_mcart_1)).
fof(t48_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>X4=k3_mcart_1(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4),k7_mcart_1(X1,X2,X3,X4)))))))), file('data/MPTP/proofs/02a8da7181ed47f983245caac44aeda0.E_input', t48_mcart_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/02a8da7181ed47f983245caac44aeda0.E_input', d3_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/02a8da7181ed47f983245caac44aeda0.E_input', d3_zfmisc_1)).
fof(t20_mcart_1, axiom, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1)))), file('data/MPTP/proofs/02a8da7181ed47f983245caac44aeda0.E_input', t20_mcart_1)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/02a8da7181ed47f983245caac44aeda0.E_input', t7_mcart_1)).
fof(t50_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))&(k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))&k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4))))))))), file('data/MPTP/proofs/02a8da7181ed47f983245caac44aeda0.E_input', t50_mcart_1)).
fof(c_0_16, plain, ![X6, X7, X8, X9, X10, X11, X12, X13]:((((r2_hidden(X9,X6)|~r2_hidden(X9,X8)|X8!=k3_xboole_0(X6,X7))&(r2_hidden(X9,X7)|~r2_hidden(X9,X8)|X8!=k3_xboole_0(X6,X7)))&(~r2_hidden(X10,X6)|~r2_hidden(X10,X7)|r2_hidden(X10,X8)|X8!=k3_xboole_0(X6,X7)))&((~r2_hidden(esk1_3(X11,X12,X13),X13)|(~r2_hidden(esk1_3(X11,X12,X13),X11)|~r2_hidden(esk1_3(X11,X12,X13),X12))|X13=k3_xboole_0(X11,X12))&((r2_hidden(esk1_3(X11,X12,X13),X11)|r2_hidden(esk1_3(X11,X12,X13),X13)|X13=k3_xboole_0(X11,X12))&(r2_hidden(esk1_3(X11,X12,X13),X12)|r2_hidden(esk1_3(X11,X12,X13),X13)|X13=k3_xboole_0(X11,X12))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_17, plain, ![X15, X16]:k4_xboole_0(X15,k4_xboole_0(X15,X16))=k3_xboole_0(X15,X16), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_18, plain, ![X17, X18, X19]:(((~r2_hidden(X17,X19)|k4_xboole_0(k2_tarski(X17,X18),X19)!=k2_tarski(X17,X18))&(~r2_hidden(X18,X19)|k4_xboole_0(k2_tarski(X17,X18),X19)!=k2_tarski(X17,X18)))&(r2_hidden(X17,X19)|r2_hidden(X18,X19)|k4_xboole_0(k2_tarski(X17,X18),X19)=k2_tarski(X17,X18))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t72_zfmisc_1])])])])).
fof(c_0_19, plain, ![X20, X21]:k1_enumset1(X20,X20,X21)=k2_tarski(X20,X21), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_20, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, plain, (~r2_hidden(X1,X2)|k4_xboole_0(k2_tarski(X3,X1),X2)!=k2_tarski(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_24, plain, ![X22]:k4_xboole_0(X22,k1_xboole_0)=X22, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_25, plain, ![X23, X24, X25, X26, X27, X28]:(((~r2_hidden(X25,X24)|X25=X23|X24!=k1_tarski(X23))&(X26!=X23|r2_hidden(X26,X24)|X24!=k1_tarski(X23)))&((~r2_hidden(esk2_2(X27,X28),X28)|esk2_2(X27,X28)!=X27|X28=k1_tarski(X27))&(r2_hidden(esk2_2(X27,X28),X28)|esk2_2(X27,X28)=X27|X28=k1_tarski(X27)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_26, plain, ![X30]:k2_tarski(X30,X30)=k1_tarski(X30), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_27, plain, (r2_hidden(X1,X2)|X3!=k4_xboole_0(X4,k4_xboole_0(X4,X2))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_28, plain, (k4_xboole_0(k1_enumset1(X3,X3,X1),X2)!=k1_enumset1(X3,X3,X1)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_23]), c_0_23])).
cnf(c_0_29, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X3,k4_xboole_0(X3,X2)))), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_33, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
fof(c_0_34, plain, ![X31, X33, X34, X35]:((r2_hidden(esk3_1(X31),X31)|X31=k1_xboole_0)&((~r2_hidden(X33,X31)|esk3_1(X31)!=k3_mcart_1(X33,X34,X35)|X31=k1_xboole_0)&(~r2_hidden(X34,X31)|esk3_1(X31)!=k3_mcart_1(X33,X34,X35)|X31=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t29_mcart_1])])])])])])])).
fof(c_0_35, plain, ![X36, X37, X38, X39, X40, X41, X42, X43, X44, X45]:(((~r2_hidden(X40,X39)|(X40=X36|(X40=X37|X40=X38))|X39!=k1_enumset1(X36,X37,X38))&((X41!=X36|r2_hidden(X41,X39)|X39!=k1_enumset1(X36,X37,X38))&((X41!=X37|r2_hidden(X41,X39)|X39!=k1_enumset1(X36,X37,X38))&(X41!=X38|r2_hidden(X41,X39)|X39!=k1_enumset1(X36,X37,X38)))))&(((esk4_4(X42,X43,X44,X45)!=X42|~r2_hidden(esk4_4(X42,X43,X44,X45),X45)|X45=k1_enumset1(X42,X43,X44))&((esk4_4(X42,X43,X44,X45)!=X43|~r2_hidden(esk4_4(X42,X43,X44,X45),X45)|X45=k1_enumset1(X42,X43,X44))&(esk4_4(X42,X43,X44,X45)!=X44|~r2_hidden(esk4_4(X42,X43,X44,X45),X45)|X45=k1_enumset1(X42,X43,X44))))&(r2_hidden(esk4_4(X42,X43,X44,X45),X45)|(esk4_4(X42,X43,X44,X45)=X42|(esk4_4(X42,X43,X44,X45)=X43|esk4_4(X42,X43,X44,X45)=X44))|X45=k1_enumset1(X42,X43,X44)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])).
fof(c_0_36, negated_conjecture, ~(![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(~(X4=k5_mcart_1(X1,X2,X3,X4))&(~(X4=k6_mcart_1(X1,X2,X3,X4))&~(X4=k7_mcart_1(X1,X2,X3,X4))))))))))), inference(assume_negation,[status(cth)],[t51_mcart_1])).
cnf(c_0_37, plain, (X1=X3|X2!=k1_enumset1(X3,X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_31]), c_0_23])).
cnf(c_0_38, plain, (~r2_hidden(X1,k4_xboole_0(X2,X2))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_29]), c_0_33])).
cnf(c_0_39, plain, (r2_hidden(esk3_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_40, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X5,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_41, plain, ![X51, X52, X53, X54]:(X51=k1_xboole_0|(X52=k1_xboole_0|(X53=k1_xboole_0|(~m1_subset_1(X54,k3_zfmisc_1(X51,X52,X53))|X54=k3_mcart_1(k5_mcart_1(X51,X52,X53,X54),k6_mcart_1(X51,X52,X53,X54),k7_mcart_1(X51,X52,X53,X54)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t48_mcart_1])])])])).
fof(c_0_42, plain, ![X58, X59, X60]:k3_mcart_1(X58,X59,X60)=k4_tarski(k4_tarski(X58,X59),X60), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_43, plain, ![X55, X56, X57]:k3_zfmisc_1(X55,X56,X57)=k2_zfmisc_1(k2_zfmisc_1(X55,X56),X57), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_44, negated_conjecture, (esk5_0!=k1_xboole_0&(esk6_0!=k1_xboole_0&(esk7_0!=k1_xboole_0&(m1_subset_1(esk8_0,k3_zfmisc_1(esk5_0,esk6_0,esk7_0))&(esk8_0=k5_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0)|(esk8_0=k6_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0)|esk8_0=k7_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_36])])])])).
fof(c_0_45, plain, ![X61, X62, X63]:((X61!=k1_mcart_1(X61)|X61!=k4_tarski(X62,X63))&(X61!=k2_mcart_1(X61)|X61!=k4_tarski(X62,X63))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t20_mcart_1])])])])])).
fof(c_0_46, plain, ![X64, X65]:(k1_mcart_1(k4_tarski(X64,X65))=X64&k2_mcart_1(k4_tarski(X64,X65))=X65), inference(variable_rename,[status(thm)],[t7_mcart_1])).
fof(c_0_47, plain, ![X66, X67, X68, X69]:((k5_mcart_1(X66,X67,X68,X69)=k1_mcart_1(k1_mcart_1(X69))|~m1_subset_1(X69,k3_zfmisc_1(X66,X67,X68))|X68=k1_xboole_0|X67=k1_xboole_0|X66=k1_xboole_0)&((k6_mcart_1(X66,X67,X68,X69)=k2_mcart_1(k1_mcart_1(X69))|~m1_subset_1(X69,k3_zfmisc_1(X66,X67,X68))|X68=k1_xboole_0|X67=k1_xboole_0|X66=k1_xboole_0)&(k7_mcart_1(X66,X67,X68,X69)=k2_mcart_1(X69)|~m1_subset_1(X69,k3_zfmisc_1(X66,X67,X68))|X68=k1_xboole_0|X67=k1_xboole_0|X66=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_mcart_1])])])])])).
cnf(c_0_48, plain, (X1=X2|~r2_hidden(X1,k1_enumset1(X2,X2,X2))), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_49, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_50, plain, (r2_hidden(X1,k1_enumset1(X2,X3,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_40])])).
cnf(c_0_51, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k3_mcart_1(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4),k7_mcart_1(X1,X2,X3,X4))|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_52, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_53, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(esk8_0,k3_zfmisc_1(esk5_0,esk6_0,esk7_0))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_55, plain, (X1!=k2_mcart_1(X1)|X1!=k4_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_56, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_57, plain, (k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_58, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk3_1(X2)!=k3_mcart_1(X3,X1,X4)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_59, plain, (esk3_1(k1_enumset1(X1,X1,X1))=X1|k1_enumset1(X1,X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_48, c_0_39])).
cnf(c_0_60, plain, (k1_enumset1(X1,X1,X2)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_49]), c_0_50])])).
cnf(c_0_61, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X4=k4_tarski(k4_tarski(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4)),k7_mcart_1(X1,X2,X3,X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_52]), c_0_53])).
cnf(c_0_62, negated_conjecture, (esk8_0=k5_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0)|esk8_0=k6_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0)|esk8_0=k7_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_63, negated_conjecture, (m1_subset_1(esk8_0,k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0))), inference(rw,[status(thm)],[c_0_54, c_0_53])).
cnf(c_0_64, plain, (k4_tarski(X1,X2)!=X2), inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_55]), c_0_56])).
cnf(c_0_65, negated_conjecture, (esk7_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_66, negated_conjecture, (esk6_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_67, negated_conjecture, (esk5_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_68, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_57, c_0_53])).
cnf(c_0_69, plain, (X2=k1_xboole_0|esk3_1(X2)!=k4_tarski(k4_tarski(X3,X1),X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_58, c_0_52])).
cnf(c_0_70, plain, (esk3_1(k1_enumset1(X1,X1,X1))=X1), inference(sr,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_71, negated_conjecture, (k6_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0)=esk8_0|k5_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0)=esk8_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63])]), c_0_64]), c_0_65]), c_0_66]), c_0_67])).
cnf(c_0_72, negated_conjecture, (k7_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0)=k2_mcart_1(esk8_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_63]), c_0_65]), c_0_66]), c_0_67])).
cnf(c_0_73, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X2,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_74, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk3_1(X2)!=k3_mcart_1(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_75, plain, (~r2_hidden(X1,k1_enumset1(k4_tarski(k4_tarski(X2,X1),X3),k4_tarski(k4_tarski(X2,X1),X3),k4_tarski(k4_tarski(X2,X1),X3)))), inference(er,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_60])])).
cnf(c_0_76, negated_conjecture, (k4_tarski(k4_tarski(k5_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0),esk8_0),k2_mcart_1(esk8_0))=esk8_0|k5_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0)=esk8_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_71]), c_0_72]), c_0_63])]), c_0_65]), c_0_66]), c_0_67])).
cnf(c_0_77, plain, (r2_hidden(X1,k1_enumset1(X1,X2,X3))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_73])])).
cnf(c_0_78, plain, (X2=k1_xboole_0|esk3_1(X2)!=k4_tarski(k4_tarski(X1,X3),X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_74, c_0_52])).
cnf(c_0_79, negated_conjecture, (k5_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0)=esk8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77])])).
cnf(c_0_80, plain, (~r2_hidden(X1,k1_enumset1(k4_tarski(k4_tarski(X1,X2),X3),k4_tarski(k4_tarski(X1,X2),X3),k4_tarski(k4_tarski(X1,X2),X3)))), inference(er,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_70]), c_0_60])])).
cnf(c_0_81, negated_conjecture, (k4_tarski(k4_tarski(esk8_0,k6_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0)),k2_mcart_1(esk8_0))=esk8_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_79]), c_0_72]), c_0_63])]), c_0_65]), c_0_66]), c_0_67])).
cnf(c_0_82, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_77])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 83
# Proof object clause steps            : 50
# Proof object formula steps           : 33
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 52
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 43
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 38
# Processed clauses                    : 244
# ...of these trivial                  : 3
# ...subsumed                          : 99
# ...remaining for further processing  : 142
# Other redundant clauses eliminated   : 53
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 6
# Generated clauses                    : 1289
# ...of the previous two non-trivial   : 1059
# Contextual simplify-reflections      : 1
# Paramodulations                      : 1235
# Factorizations                       : 4
# Equation resolutions                 : 53
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 83
#    Positive orientable unit clauses  : 20
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 10
#    Non-unit-clauses                  : 53
# Current number of unprocessed clauses: 857
# ...number of literals in the above   : 3118
# Current number of archived formulas  : 0
# Current number of archived clauses   : 53
# Clause-clause subsumption calls (NU) : 409
# Rec. Clause-clause subsumption calls : 310
# Non-unit clause-clause subsumptions  : 70
# Unit Clause-clause subsumption calls : 68
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 18
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 17564

# -------------------------------------------------
# User time                : 0.018 s
# System time              : 0.003 s
# Total time               : 0.021 s
# Maximum resident set size: 28684 pages
