# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d20_borsuk_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>(r1_borsuk_1(X1,X2)<=>?[X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&(v5_pre_topc(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))))&v3_borsuk_1(X3,X1,X2))))), file('data/MPTP/proofs/02ddc04aa6f44498aa4815df55d0ae32.E_input', d20_borsuk_1)).
fof(t78_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v3_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:((~(v2_struct_0(X2))&(v1_tdlat_3(X2)&m1_pre_topc(X2,X1)))=>?[X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&(v5_pre_topc(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))))&v3_borsuk_1(X3,X1,X2)))), file('data/MPTP/proofs/02ddc04aa6f44498aa4815df55d0ae32.E_input', t78_tex_2)).
fof(t79_tex_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v3_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:((~(v2_struct_0(X2))&(v1_tdlat_3(X2)&m1_pre_topc(X2,X1)))=>r1_borsuk_1(X1,X2))), file('data/MPTP/proofs/02ddc04aa6f44498aa4815df55d0ae32.E_input', t79_tex_2)).
fof(c_0_3, plain, ![X6, X7, X9]:((((v1_funct_1(esk3_2(X6,X7))|~r1_borsuk_1(X6,X7)|(v2_struct_0(X7)|~m1_pre_topc(X7,X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&((v1_funct_2(esk3_2(X6,X7),u1_struct_0(X6),u1_struct_0(X7))|~r1_borsuk_1(X6,X7)|(v2_struct_0(X7)|~m1_pre_topc(X7,X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&((v5_pre_topc(esk3_2(X6,X7),X6,X7)|~r1_borsuk_1(X6,X7)|(v2_struct_0(X7)|~m1_pre_topc(X7,X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&(m1_subset_1(esk3_2(X6,X7),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X6),u1_struct_0(X7))))|~r1_borsuk_1(X6,X7)|(v2_struct_0(X7)|~m1_pre_topc(X7,X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6)))))))&(v3_borsuk_1(esk3_2(X6,X7),X6,X7)|~r1_borsuk_1(X6,X7)|(v2_struct_0(X7)|~m1_pre_topc(X7,X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6)))))&(~v1_funct_1(X9)|(~v1_funct_2(X9,u1_struct_0(X6),u1_struct_0(X7))|(~v5_pre_topc(X9,X6,X7)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X6),u1_struct_0(X7))))))|~v3_borsuk_1(X9,X6,X7)|r1_borsuk_1(X6,X7)|(v2_struct_0(X7)|~m1_pre_topc(X7,X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d20_borsuk_1])])])])])])).
fof(c_0_4, plain, ![X10, X11]:(((v1_funct_1(esk4_2(X10,X11))|(v2_struct_0(X11)|(~v1_tdlat_3(X11)|~m1_pre_topc(X11,X10)))|(v2_struct_0(X10)|(~v2_pre_topc(X10)|(~v3_tdlat_3(X10)|~l1_pre_topc(X10)))))&((v1_funct_2(esk4_2(X10,X11),u1_struct_0(X10),u1_struct_0(X11))|(v2_struct_0(X11)|(~v1_tdlat_3(X11)|~m1_pre_topc(X11,X10)))|(v2_struct_0(X10)|(~v2_pre_topc(X10)|(~v3_tdlat_3(X10)|~l1_pre_topc(X10)))))&((v5_pre_topc(esk4_2(X10,X11),X10,X11)|(v2_struct_0(X11)|(~v1_tdlat_3(X11)|~m1_pre_topc(X11,X10)))|(v2_struct_0(X10)|(~v2_pre_topc(X10)|(~v3_tdlat_3(X10)|~l1_pre_topc(X10)))))&(m1_subset_1(esk4_2(X10,X11),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X11))))|(v2_struct_0(X11)|(~v1_tdlat_3(X11)|~m1_pre_topc(X11,X10)))|(v2_struct_0(X10)|(~v2_pre_topc(X10)|(~v3_tdlat_3(X10)|~l1_pre_topc(X10))))))))&(v3_borsuk_1(esk4_2(X10,X11),X10,X11)|(v2_struct_0(X11)|(~v1_tdlat_3(X11)|~m1_pre_topc(X11,X10)))|(v2_struct_0(X10)|(~v2_pre_topc(X10)|(~v3_tdlat_3(X10)|~l1_pre_topc(X10)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t78_tex_2])])])])])])).
fof(c_0_5, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v3_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:((~(v2_struct_0(X2))&(v1_tdlat_3(X2)&m1_pre_topc(X2,X1)))=>r1_borsuk_1(X1,X2)))), inference(assume_negation,[status(cth)],[t79_tex_2])).
cnf(c_0_6, plain, (r1_borsuk_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v5_pre_topc(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v3_borsuk_1(X1,X2,X3)|~m1_pre_topc(X3,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_7, plain, (m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_tdlat_3(X2)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~v3_tdlat_3(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, plain, (v1_funct_1(esk4_2(X1,X2))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_tdlat_3(X2)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~v3_tdlat_3(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9, plain, (v1_funct_2(esk4_2(X1,X2),u1_struct_0(X1),u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_tdlat_3(X2)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~v3_tdlat_3(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10, plain, (v5_pre_topc(esk4_2(X1,X2),X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|~v1_tdlat_3(X2)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~v3_tdlat_3(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_11, plain, (v3_borsuk_1(esk4_2(X1,X2),X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|~v1_tdlat_3(X2)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~v3_tdlat_3(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_12, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&(v3_tdlat_3(esk1_0)&l1_pre_topc(esk1_0))))&((~v2_struct_0(esk2_0)&(v1_tdlat_3(esk2_0)&m1_pre_topc(esk2_0,esk1_0)))&~r1_borsuk_1(esk1_0,esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_5])])])])).
cnf(c_0_13, plain, (r1_borsuk_1(X1,X2)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_pre_topc(X2,X1)|~v1_tdlat_3(X2)|~l1_pre_topc(X1)|~v3_tdlat_3(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8]), c_0_9]), c_0_10]), c_0_11])).
cnf(c_0_14, negated_conjecture, (v1_tdlat_3(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, negated_conjecture, (r1_borsuk_1(X1,esk2_0)|v2_struct_0(X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)|~v3_tdlat_3(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])).
cnf(c_0_17, negated_conjecture, (v3_tdlat_3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, negated_conjecture, (~r1_borsuk_1(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18]), c_0_19]), c_0_20])]), c_0_21]), c_0_22]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 24
# Proof object clause steps            : 17
# Proof object formula steps           : 7
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 19
# Processed clauses                    : 40
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 40
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 4
# ...of the previous two non-trivial   : 2
# Contextual simplify-reflections      : 4
# Paramodulations                      : 4
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 21
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 13
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 19
# Clause-clause subsumption calls (NU) : 35
# Rec. Clause-clause subsumption calls : 4
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2043

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28624 pages
