# No SInE strategy applied
# Auto-Mode selected heuristic G_____0017_C18_F1_SE_CS_SP_S4Y
# and selection function SelectMaxLComplexAPPNTNp.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/037f92227da14aceb21e4d3b4f856f73.E_input', t17_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/037f92227da14aceb21e4d3b4f856f73.E_input', t48_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/037f92227da14aceb21e4d3b4f856f73.E_input', commutativity_k3_xboole_0)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/037f92227da14aceb21e4d3b4f856f73.E_input', d7_xboole_0)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/037f92227da14aceb21e4d3b4f856f73.E_input', t3_boole)).
fof(t63_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_xboole_0(X2,X3))=>r1_xboole_0(X1,X3)), file('data/MPTP/proofs/037f92227da14aceb21e4d3b4f856f73.E_input', t63_xboole_1)).
fof(t70_xboole_1, axiom, ![X1, X2, X3]:(~((~(r1_xboole_0(X1,k2_xboole_0(X2,X3)))&(r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3))))&~((~((r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3)))&r1_xboole_0(X1,k2_xboole_0(X2,X3))))), file('data/MPTP/proofs/037f92227da14aceb21e4d3b4f856f73.E_input', t70_xboole_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/037f92227da14aceb21e4d3b4f856f73.E_input', t12_xboole_1)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/037f92227da14aceb21e4d3b4f856f73.E_input', symmetry_r1_xboole_0)).
fof(t77_xboole_1, conjecture, ![X1, X2, X3]:~((~(r1_xboole_0(X1,X2))&(r1_tarski(X1,X3)&r1_xboole_0(X1,k3_xboole_0(X2,X3))))), file('data/MPTP/proofs/037f92227da14aceb21e4d3b4f856f73.E_input', t77_xboole_1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/037f92227da14aceb21e4d3b4f856f73.E_input', t4_xboole_0)).
fof(t26_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(k3_xboole_0(X1,X3),k3_xboole_0(X2,X3))), file('data/MPTP/proofs/037f92227da14aceb21e4d3b4f856f73.E_input', t26_xboole_1)).
fof(c_0_12, plain, ![X4, X5]:r1_tarski(k3_xboole_0(X4,X5),X4), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_13, plain, ![X6, X7]:k4_xboole_0(X6,k4_xboole_0(X6,X7))=k3_xboole_0(X6,X7), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_14, plain, ![X8, X9]:k3_xboole_0(X8,X9)=k3_xboole_0(X9,X8), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_15, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, plain, ![X10, X11]:((~r1_xboole_0(X10,X11)|k3_xboole_0(X10,X11)=k1_xboole_0)&(k3_xboole_0(X10,X11)!=k1_xboole_0|r1_xboole_0(X10,X11))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_19, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X1)), inference(rw,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_20, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_16]), c_0_16])).
fof(c_0_21, plain, ![X12]:k4_xboole_0(X12,k1_xboole_0)=X12, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_22, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_23, plain, ![X13, X14, X15]:(~r1_tarski(X13,X14)|~r1_xboole_0(X14,X15)|r1_xboole_0(X13,X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_xboole_1])])).
cnf(c_0_24, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_25, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_26, plain, ![X16, X17, X18]:((r1_xboole_0(X16,k2_xboole_0(X17,X18))|(~r1_xboole_0(X16,X17)|~r1_xboole_0(X16,X18)))&((r1_xboole_0(X16,X17)|~r1_xboole_0(X16,k2_xboole_0(X17,X18)))&(r1_xboole_0(X16,X18)|~r1_xboole_0(X16,k2_xboole_0(X17,X18))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t70_xboole_1])])])])).
fof(c_0_27, plain, ![X19, X20]:(~r1_tarski(X19,X20)|k2_xboole_0(X19,X20)=X20), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_28, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_29, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_22, c_0_16])).
cnf(c_0_30, plain, (r1_xboole_0(X1,X3)|~r1_tarski(X1,X2)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31, plain, (r1_tarski(k4_xboole_0(X1,X1),k1_xboole_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_32, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_34, plain, (r1_xboole_0(X1,X2)|k4_xboole_0(X1,k4_xboole_0(X1,X2))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_28, c_0_16])).
cnf(c_0_35, plain, (r1_tarski(k1_xboole_0,X1)|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_19, c_0_29])).
cnf(c_0_36, plain, (r1_xboole_0(k4_xboole_0(X1,X1),X2)|~r1_xboole_0(k1_xboole_0,X2)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_37, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X1,X3)|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_38, plain, (r1_xboole_0(X1,k1_xboole_0)|k4_xboole_0(X1,X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_34, c_0_25])).
cnf(c_0_39, plain, (r1_tarski(k1_xboole_0,k4_xboole_0(X1,X1))|~r1_xboole_0(k1_xboole_0,X2)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_40, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X2))|~r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_37, c_0_31])).
cnf(c_0_41, plain, (r1_xboole_0(k1_xboole_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_38, c_0_25])).
cnf(c_0_42, plain, (r1_tarski(k1_xboole_0,k4_xboole_0(X1,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])])).
cnf(c_0_43, plain, (r1_tarski(k4_xboole_0(X1,X1),X1)), inference(spm,[status(thm)],[c_0_19, c_0_25])).
fof(c_0_44, plain, ![X21, X22]:(~r1_xboole_0(X21,X22)|r1_xboole_0(X22,X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
fof(c_0_45, negated_conjecture, ~(![X1, X2, X3]:~((~(r1_xboole_0(X1,X2))&(r1_tarski(X1,X3)&r1_xboole_0(X1,k3_xboole_0(X2,X3)))))), inference(assume_negation,[status(cth)],[t77_xboole_1])).
cnf(c_0_46, plain, (r1_xboole_0(k1_xboole_0,X1)|~r1_xboole_0(k4_xboole_0(X2,X2),X1)), inference(spm,[status(thm)],[c_0_30, c_0_42])).
cnf(c_0_47, plain, (r1_xboole_0(k4_xboole_0(X1,X1),X2)|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_30, c_0_43])).
cnf(c_0_48, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_49, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X2))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_37, c_0_43])).
fof(c_0_50, negated_conjecture, (~r1_xboole_0(esk1_0,esk2_0)&(r1_tarski(esk1_0,esk3_0)&r1_xboole_0(esk1_0,k3_xboole_0(esk2_0,esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_45])])])])).
cnf(c_0_51, plain, (r1_xboole_0(k1_xboole_0,X1)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_52, plain, (r1_xboole_0(k4_xboole_0(X1,X1),X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_53, negated_conjecture, (r1_xboole_0(esk1_0,k3_xboole_0(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_54, plain, (r1_tarski(X1,k4_xboole_0(X1,X2))|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_29]), c_0_25])).
cnf(c_0_55, plain, (r1_xboole_0(k1_xboole_0,X1)|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_56, negated_conjecture, (r1_xboole_0(esk1_0,k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0)))), inference(rw,[status(thm)],[c_0_53, c_0_16])).
cnf(c_0_57, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(k4_xboole_0(X1,X3),X2)|~r1_xboole_0(X1,X3)), inference(spm,[status(thm)],[c_0_30, c_0_54])).
cnf(c_0_58, negated_conjecture, (r1_xboole_0(k1_xboole_0,esk1_0)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_59, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(k1_xboole_0,X2)|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_57, c_0_36])).
cnf(c_0_60, plain, (r1_xboole_0(k1_xboole_0,k4_xboole_0(X1,X1))|~r1_xboole_0(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_51, c_0_40])).
cnf(c_0_61, negated_conjecture, (r1_xboole_0(esk1_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_48, c_0_58])).
fof(c_0_62, plain, ![X26, X27, X29, X30, X31]:((r1_xboole_0(X26,X27)|r2_hidden(esk4_2(X26,X27),k3_xboole_0(X26,X27)))&(~r2_hidden(X31,k3_xboole_0(X29,X30))|~r1_xboole_0(X29,X30))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
cnf(c_0_63, plain, (X1=k1_xboole_0|~r1_xboole_0(X1,k4_xboole_0(X1,X2))|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_29]), c_0_25])).
cnf(c_0_64, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X2))|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_48, c_0_47])).
cnf(c_0_65, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X2))|~r1_xboole_0(X1,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_40]), c_0_41])])).
cnf(c_0_66, negated_conjecture, (r1_xboole_0(k1_xboole_0,k4_xboole_0(X1,X1))), inference(spm,[status(thm)],[c_0_60, c_0_61])).
fof(c_0_67, plain, ![X32, X33, X34]:(~r1_tarski(X32,X33)|r1_tarski(k3_xboole_0(X32,X34),k3_xboole_0(X33,X34))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t26_xboole_1])])).
cnf(c_0_68, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_69, plain, (X1=k1_xboole_0|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_70, plain, (r1_xboole_0(k4_xboole_0(X1,X1),k4_xboole_0(X2,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_36]), c_0_66])])).
cnf(c_0_71, plain, (r1_tarski(k3_xboole_0(X1,X3),k3_xboole_0(X2,X3))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_72, plain, (~r1_xboole_0(X2,X3)|~r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))), inference(rw,[status(thm)],[c_0_68, c_0_16])).
cnf(c_0_73, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_74, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk4_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_75, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X3)),k4_xboole_0(X2,k4_xboole_0(X2,X3)))|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_16]), c_0_16])).
cnf(c_0_76, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(X2,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_25])).
cnf(c_0_77, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk4_2(X1,X2),k4_xboole_0(X1,k4_xboole_0(X1,X2)))), inference(rw,[status(thm)],[c_0_74, c_0_16])).
cnf(c_0_78, plain, (r1_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X3)|~r1_xboole_0(k4_xboole_0(X4,k4_xboole_0(X4,X2)),X3)|~r1_tarski(X1,X4)), inference(spm,[status(thm)],[c_0_30, c_0_75])).
cnf(c_0_79, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k4_xboole_0(X1,k4_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_80, plain, (r1_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_37, c_0_19])).
cnf(c_0_81, plain, (r1_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X3)|~r1_xboole_0(k4_xboole_0(X2,k4_xboole_0(X2,X4)),X3)|~r1_tarski(X1,X4)), inference(spm,[status(thm)],[c_0_78, c_0_20])).
cnf(c_0_82, negated_conjecture, (r1_xboole_0(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0)),esk1_0)), inference(spm,[status(thm)],[c_0_48, c_0_56])).
cnf(c_0_83, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_84, negated_conjecture, (r1_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,esk2_0)),esk1_0)|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_85, negated_conjecture, (r1_tarski(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_86, negated_conjecture, (~r1_xboole_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_87, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_85])]), c_0_86]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 88
# Proof object clause steps            : 63
# Proof object formula steps           : 25
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 39
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 17
# Processed clauses                    : 9784
# ...of these trivial                  : 95
# ...subsumed                          : 8759
# ...remaining for further processing  : 930
# Other redundant clauses eliminated   : 30
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 33
# Backward-rewritten                   : 81
# Generated clauses                    : 47998
# ...of the previous two non-trivial   : 40690
# Contextual simplify-reflections      : 3
# Paramodulations                      : 47968
# Factorizations                       : 0
# Equation resolutions                 : 30
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 816
#    Positive orientable unit clauses  : 20
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 19
#    Non-unit-clauses                  : 776
# Current number of unprocessed clauses: 30310
# ...number of literals in the above   : 90420
# Current number of archived formulas  : 0
# Current number of archived clauses   : 115
# Clause-clause subsumption calls (NU) : 236433
# Rec. Clause-clause subsumption calls : 198837
# Non-unit clause-clause subsumptions  : 6663
# Unit Clause-clause subsumption calls : 1541
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 105
# BW rewrite match successes           : 41
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 656999

# -------------------------------------------------
# User time                : 0.266 s
# System time              : 0.007 s
# Total time               : 0.272 s
# Maximum resident set size: 28680 pages
