# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/03c594e9cf9c4d13921d1f8cdcd744dd.E_input', d2_subset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/03c594e9cf9c4d13921d1f8cdcd744dd.E_input', t3_subset)).
fof(t44_tex_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v2_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v2_tex_2(X2,X1)<=>v1_zfmisc_1(X2)))), file('data/MPTP/proofs/03c594e9cf9c4d13921d1f8cdcd744dd.E_input', t44_tex_2)).
fof(l1_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/03c594e9cf9c4d13921d1f8cdcd744dd.E_input', l1_zfmisc_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/03c594e9cf9c4d13921d1f8cdcd744dd.E_input', d1_xboole_0)).
fof(t1_tex_2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((~(v1_xboole_0(X2))&v1_zfmisc_1(X2))=>(r1_tarski(X1,X2)=>X1=X2))), file('data/MPTP/proofs/03c594e9cf9c4d13921d1f8cdcd744dd.E_input', t1_tex_2)).
fof(fc2_xboole_0, axiom, ![X1]:~(v1_xboole_0(k1_tarski(X1))), file('data/MPTP/proofs/03c594e9cf9c4d13921d1f8cdcd744dd.E_input', fc2_xboole_0)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/03c594e9cf9c4d13921d1f8cdcd744dd.E_input', redefinition_k6_domain_1)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/03c594e9cf9c4d13921d1f8cdcd744dd.E_input', t1_xboole_1)).
fof(t36_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>v2_tex_2(k6_domain_1(u1_struct_0(X1),X2),X1))), file('data/MPTP/proofs/03c594e9cf9c4d13921d1f8cdcd744dd.E_input', t36_tex_2)).
fof(t34_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>~((v2_tex_2(X2,X1)&![X3]:((~(v2_struct_0(X3))&(v1_pre_topc(X3)&(v1_tdlat_3(X3)&m1_pre_topc(X3,X1))))=>~(X2=u1_struct_0(X3))))))), file('data/MPTP/proofs/03c594e9cf9c4d13921d1f8cdcd744dd.E_input', t34_tex_2)).
fof(cc6_tdlat_3, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v2_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:(m1_pre_topc(X2,X1)=>v2_tdlat_3(X2))), file('data/MPTP/proofs/03c594e9cf9c4d13921d1f8cdcd744dd.E_input', cc6_tdlat_3)).
fof(cc2_tdlat_3, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_tdlat_3(X1)=>v2_pre_topc(X1))), file('data/MPTP/proofs/03c594e9cf9c4d13921d1f8cdcd744dd.E_input', cc2_tdlat_3)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/03c594e9cf9c4d13921d1f8cdcd744dd.E_input', dt_m1_pre_topc)).
fof(cc2_tex_1, axiom, ![X1]:(l1_pre_topc(X1)=>((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v1_tdlat_3(X1)&v2_tdlat_3(X1))))=>(~(v2_struct_0(X1))&(v7_struct_0(X1)&v2_pre_topc(X1))))), file('data/MPTP/proofs/03c594e9cf9c4d13921d1f8cdcd744dd.E_input', cc2_tex_1)).
fof(cc1_tdlat_3, axiom, ![X1]:(l1_pre_topc(X1)=>(v1_tdlat_3(X1)=>v2_pre_topc(X1))), file('data/MPTP/proofs/03c594e9cf9c4d13921d1f8cdcd744dd.E_input', cc1_tdlat_3)).
fof(fc7_struct_0, axiom, ![X1]:((v7_struct_0(X1)&l1_struct_0(X1))=>v1_zfmisc_1(u1_struct_0(X1))), file('data/MPTP/proofs/03c594e9cf9c4d13921d1f8cdcd744dd.E_input', fc7_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/03c594e9cf9c4d13921d1f8cdcd744dd.E_input', dt_l1_pre_topc)).
fof(cc1_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>(v2_struct_0(X1)=>v7_struct_0(X1))), file('data/MPTP/proofs/03c594e9cf9c4d13921d1f8cdcd744dd.E_input', cc1_struct_0)).
fof(c_0_19, plain, ![X29, X30]:(((~m1_subset_1(X30,X29)|r2_hidden(X30,X29)|v1_xboole_0(X29))&(~r2_hidden(X30,X29)|m1_subset_1(X30,X29)|v1_xboole_0(X29)))&((~m1_subset_1(X30,X29)|v1_xboole_0(X30)|~v1_xboole_0(X29))&(~v1_xboole_0(X30)|m1_subset_1(X30,X29)|~v1_xboole_0(X29)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_20, plain, ![X37, X38]:((~m1_subset_1(X37,k1_zfmisc_1(X38))|r1_tarski(X37,X38))&(~r1_tarski(X37,X38)|m1_subset_1(X37,k1_zfmisc_1(X38)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_21, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v2_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v2_tex_2(X2,X1)<=>v1_zfmisc_1(X2))))), inference(assume_negation,[status(cth)],[t44_tex_2])).
cnf(c_0_22, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_24, plain, ![X10, X11]:((~r1_tarski(k1_tarski(X10),X11)|r2_hidden(X10,X11))&(~r2_hidden(X10,X11)|r1_tarski(k1_tarski(X10),X11))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])).
fof(c_0_25, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&(v2_tdlat_3(esk1_0)&l1_pre_topc(esk1_0))))&((~v1_xboole_0(esk2_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))&((~v2_tex_2(esk2_0,esk1_0)|~v1_zfmisc_1(esk2_0))&(v2_tex_2(esk2_0,esk1_0)|v1_zfmisc_1(esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_21])])])])).
fof(c_0_26, plain, ![X18, X19, X20]:((~v1_xboole_0(X18)|~r2_hidden(X19,X18))&(r2_hidden(esk3_1(X20),X20)|v1_xboole_0(X20))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_27, plain, (r2_hidden(X1,k1_zfmisc_1(X2))|v1_xboole_0(k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_28, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30, plain, (r2_hidden(esk3_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_31, plain, ![X24, X25]:(v1_xboole_0(X24)|(v1_xboole_0(X25)|~v1_zfmisc_1(X25)|(~r1_tarski(X24,X25)|X24=X25))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t1_tex_2])])])])).
fof(c_0_32, plain, ![X12]:~v1_xboole_0(k1_tarski(X12)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_xboole_0])])).
cnf(c_0_33, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_34, plain, (r2_hidden(k1_tarski(X1),k1_zfmisc_1(X2))|v1_xboole_0(k1_zfmisc_1(X2))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_35, negated_conjecture, (r2_hidden(esk3_1(esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
fof(c_0_36, plain, ![X13, X14]:(v1_xboole_0(X13)|~m1_subset_1(X14,X13)|k6_domain_1(X13,X14)=k1_tarski(X14)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_37, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_38, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_39, plain, ![X15, X16, X17]:(~r1_tarski(X15,X16)|~r1_tarski(X16,X17)|r1_tarski(X15,X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
cnf(c_0_40, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_41, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_42, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|X1=X2|~v1_zfmisc_1(X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_43, plain, (~v1_xboole_0(k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_44, plain, (v1_xboole_0(X1)|~r1_tarski(X1,X2)|~v1_xboole_0(k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_33, c_0_23])).
cnf(c_0_45, negated_conjecture, (r2_hidden(k1_tarski(esk3_1(esk2_0)),k1_zfmisc_1(esk2_0))|v1_xboole_0(k1_zfmisc_1(esk2_0))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
fof(c_0_46, plain, ![X22, X23]:(v2_struct_0(X22)|(~v2_pre_topc(X22)|~l1_pre_topc(X22))|(~m1_subset_1(X23,u1_struct_0(X22))|v2_tex_2(k6_domain_1(u1_struct_0(X22),X23),X22))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t36_tex_2])])])])).
cnf(c_0_47, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_48, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_49, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_50, negated_conjecture, (r1_tarski(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
fof(c_0_51, plain, ![X31, X32]:(((~v2_struct_0(esk4_2(X31,X32))|~v2_tex_2(X32,X31)|(v1_xboole_0(X32)|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31))))|(v2_struct_0(X31)|(~v2_pre_topc(X31)|~l1_pre_topc(X31))))&((v1_pre_topc(esk4_2(X31,X32))|~v2_tex_2(X32,X31)|(v1_xboole_0(X32)|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31))))|(v2_struct_0(X31)|(~v2_pre_topc(X31)|~l1_pre_topc(X31))))&((v1_tdlat_3(esk4_2(X31,X32))|~v2_tex_2(X32,X31)|(v1_xboole_0(X32)|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31))))|(v2_struct_0(X31)|(~v2_pre_topc(X31)|~l1_pre_topc(X31))))&(m1_pre_topc(esk4_2(X31,X32),X31)|~v2_tex_2(X32,X31)|(v1_xboole_0(X32)|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31))))|(v2_struct_0(X31)|(~v2_pre_topc(X31)|~l1_pre_topc(X31)))))))&(X32=u1_struct_0(esk4_2(X31,X32))|~v2_tex_2(X32,X31)|(v1_xboole_0(X32)|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31))))|(v2_struct_0(X31)|(~v2_pre_topc(X31)|~l1_pre_topc(X31))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_tex_2])])])])])])).
cnf(c_0_52, plain, (k1_tarski(X1)=X2|~r2_hidden(X1,X2)|~v1_zfmisc_1(X2)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_28]), c_0_43]), c_0_38])).
cnf(c_0_53, negated_conjecture, (r2_hidden(k1_tarski(esk3_1(esk2_0)),k1_zfmisc_1(esk2_0))|v1_xboole_0(X1)|~r1_tarski(X1,esk2_0)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_54, plain, (v2_struct_0(X1)|v2_tex_2(k6_domain_1(u1_struct_0(X1),X2),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_55, plain, (k6_domain_1(X1,X2)=k1_tarski(X2)|~r2_hidden(X2,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_38])).
cnf(c_0_56, plain, (r2_hidden(X1,X2)|~r1_tarski(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_57, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~r1_tarski(X1,esk2_0)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_58, plain, (m1_pre_topc(esk4_2(X1,X2),X1)|v1_xboole_0(X2)|v2_struct_0(X1)|~v2_tex_2(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_59, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_60, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_61, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_62, negated_conjecture, (k1_tarski(esk3_1(esk2_0))=esk2_0|~v1_zfmisc_1(esk2_0)), inference(spm,[status(thm)],[c_0_52, c_0_35])).
cnf(c_0_63, negated_conjecture, (v2_tex_2(esk2_0,esk1_0)|v1_zfmisc_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_64, negated_conjecture, (r2_hidden(k1_tarski(esk3_1(esk2_0)),k1_zfmisc_1(esk2_0))|~r2_hidden(X1,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_28]), c_0_43])).
cnf(c_0_65, plain, (v2_tex_2(k1_tarski(X1),X2)|v2_struct_0(X2)|~r2_hidden(X1,u1_struct_0(X2))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_48])).
cnf(c_0_66, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r1_tarski(k1_tarski(X1),esk2_0)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_67, negated_conjecture, (m1_pre_topc(esk4_2(esk1_0,esk2_0),esk1_0)|~v2_tex_2(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_41]), c_0_59]), c_0_60])]), c_0_29]), c_0_61])).
cnf(c_0_68, negated_conjecture, (k1_tarski(esk3_1(esk2_0))=esk2_0|v2_tex_2(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_69, plain, (v1_tdlat_3(esk4_2(X1,X2))|v1_xboole_0(X2)|v2_struct_0(X1)|~v2_tex_2(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_70, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_40, c_0_48])).
cnf(c_0_71, negated_conjecture, (r2_hidden(k1_tarski(esk3_1(esk2_0)),k1_zfmisc_1(esk2_0))), inference(spm,[status(thm)],[c_0_64, c_0_35])).
cnf(c_0_72, negated_conjecture, (v2_tex_2(k1_tarski(X1),esk1_0)|~r1_tarski(k1_tarski(X1),esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_59]), c_0_60])]), c_0_61])).
cnf(c_0_73, negated_conjecture, (k1_tarski(esk3_1(esk2_0))=esk2_0|m1_pre_topc(esk4_2(esk1_0,esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
fof(c_0_74, plain, ![X8, X9]:(v2_struct_0(X8)|(~v2_pre_topc(X8)|(~v2_tdlat_3(X8)|~l1_pre_topc(X8)))|(~m1_pre_topc(X9,X8)|v2_tdlat_3(X9))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc6_tdlat_3])])])])).
fof(c_0_75, plain, ![X34]:(~l1_pre_topc(X34)|(~v2_tdlat_3(X34)|v2_pre_topc(X34))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_tdlat_3])])).
cnf(c_0_76, negated_conjecture, (v1_tdlat_3(esk4_2(esk1_0,esk2_0))|~v2_tex_2(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_41]), c_0_59]), c_0_60])]), c_0_29]), c_0_61])).
fof(c_0_77, plain, ![X26, X27]:(~l1_pre_topc(X26)|(~m1_pre_topc(X27,X26)|l1_pre_topc(X27))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_78, negated_conjecture, (r1_tarski(k1_tarski(esk3_1(esk2_0)),esk2_0)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_79, negated_conjecture, (m1_pre_topc(esk4_2(esk1_0,esk2_0),esk1_0)|~r1_tarski(esk2_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_67])).
fof(c_0_80, plain, ![X28]:((~v2_struct_0(X28)|(v2_struct_0(X28)|(~v2_pre_topc(X28)|(~v1_tdlat_3(X28)|~v2_tdlat_3(X28))))|~l1_pre_topc(X28))&((v7_struct_0(X28)|(v2_struct_0(X28)|(~v2_pre_topc(X28)|(~v1_tdlat_3(X28)|~v2_tdlat_3(X28))))|~l1_pre_topc(X28))&(v2_pre_topc(X28)|(v2_struct_0(X28)|(~v2_pre_topc(X28)|(~v1_tdlat_3(X28)|~v2_tdlat_3(X28))))|~l1_pre_topc(X28)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_tex_1])])])])).
fof(c_0_81, plain, ![X7]:(~l1_pre_topc(X7)|(~v1_tdlat_3(X7)|v2_pre_topc(X7))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_tdlat_3])])).
cnf(c_0_82, plain, (v2_struct_0(X1)|v2_tdlat_3(X2)|~v2_pre_topc(X1)|~v2_tdlat_3(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_83, plain, (v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_tdlat_3(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_84, negated_conjecture, (k1_tarski(esk3_1(esk2_0))=esk2_0|v1_tdlat_3(esk4_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_76, c_0_68])).
fof(c_0_85, plain, ![X35]:(~v7_struct_0(X35)|~l1_struct_0(X35)|v1_zfmisc_1(u1_struct_0(X35))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc7_struct_0])])).
cnf(c_0_86, plain, (X1=u1_struct_0(esk4_2(X2,X1))|v1_xboole_0(X1)|v2_struct_0(X2)|~v2_tex_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
fof(c_0_87, plain, ![X6]:(~l1_pre_topc(X6)|l1_struct_0(X6)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_88, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_89, negated_conjecture, (m1_pre_topc(esk4_2(esk1_0,esk2_0),esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_73]), c_0_79])).
cnf(c_0_90, plain, (v7_struct_0(X1)|v2_struct_0(X1)|~v2_pre_topc(X1)|~v1_tdlat_3(X1)|~v2_tdlat_3(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_91, plain, (v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_92, plain, (v2_tdlat_3(X1)|v2_struct_0(X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_tdlat_3(X2)), inference(csr,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_93, negated_conjecture, (v2_tdlat_3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_94, negated_conjecture, (v1_tdlat_3(esk4_2(esk1_0,esk2_0))|~r1_tarski(esk2_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_84]), c_0_76])).
cnf(c_0_95, plain, (v1_zfmisc_1(u1_struct_0(X1))|~v7_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_96, negated_conjecture, (u1_struct_0(esk4_2(esk1_0,esk2_0))=esk2_0|~v2_tex_2(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_41]), c_0_59]), c_0_60])]), c_0_29]), c_0_61])).
cnf(c_0_97, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_98, negated_conjecture, (l1_pre_topc(esk4_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_59])])).
fof(c_0_99, plain, ![X36]:(~l1_struct_0(X36)|(~v2_struct_0(X36)|v7_struct_0(X36))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_struct_0])])).
cnf(c_0_100, plain, (v7_struct_0(X1)|v2_struct_0(X1)|~v1_tdlat_3(X1)|~l1_pre_topc(X1)|~v2_tdlat_3(X1)), inference(csr,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_101, negated_conjecture, (v2_tdlat_3(esk4_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_89]), c_0_59]), c_0_93])]), c_0_61])).
cnf(c_0_102, negated_conjecture, (v1_tdlat_3(esk4_2(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_84]), c_0_94])).
cnf(c_0_103, negated_conjecture, (v1_zfmisc_1(esk2_0)|~v7_struct_0(esk4_2(esk1_0,esk2_0))|~l1_struct_0(esk4_2(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_63])).
cnf(c_0_104, negated_conjecture, (l1_struct_0(esk4_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_105, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_106, negated_conjecture, (v7_struct_0(esk4_2(esk1_0,esk2_0))|v2_struct_0(esk4_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_102]), c_0_98])])).
cnf(c_0_107, negated_conjecture, (v1_zfmisc_1(esk2_0)|~v7_struct_0(esk4_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_104])])).
cnf(c_0_108, negated_conjecture, (v7_struct_0(esk4_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_104])])).
cnf(c_0_109, negated_conjecture, (~v2_tex_2(esk2_0,esk1_0)|~v1_zfmisc_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_110, negated_conjecture, (v1_zfmisc_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_107, c_0_108])])).
cnf(c_0_111, negated_conjecture, (~v7_struct_0(esk4_2(esk1_0,esk2_0))|~l1_struct_0(esk4_2(esk1_0,esk2_0))|~v2_tex_2(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_109, c_0_103])).
cnf(c_0_112, negated_conjecture, (k1_tarski(esk3_1(esk2_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_110])])).
cnf(c_0_113, negated_conjecture, (~v7_struct_0(esk4_2(esk1_0,esk2_0))|~v2_tex_2(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_111, c_0_104])])).
cnf(c_0_114, negated_conjecture, (r1_tarski(esk2_0,esk2_0)), inference(rw,[status(thm)],[c_0_78, c_0_112])).
cnf(c_0_115, negated_conjecture, (~v2_tex_2(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_113, c_0_108])])).
cnf(c_0_116, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_112]), c_0_114])]), c_0_115]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 117
# Proof object clause steps            : 78
# Proof object formula steps           : 39
# Proof object conjectures             : 46
# Proof object clause conjectures      : 43
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 33
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 36
# Proof object simplifying inferences  : 57
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 38
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 36
# Processed clauses                    : 331
# ...of these trivial                  : 3
# ...subsumed                          : 88
# ...remaining for further processing  : 239
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 17
# Backward-rewritten                   : 32
# Generated clauses                    : 601
# ...of the previous two non-trivial   : 553
# Contextual simplify-reflections      : 23
# Paramodulations                      : 601
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 154
#    Positive orientable unit clauses  : 27
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 8
#    Non-unit-clauses                  : 119
# Current number of unprocessed clauses: 286
# ...number of literals in the above   : 1333
# Current number of archived formulas  : 0
# Current number of archived clauses   : 85
# Clause-clause subsumption calls (NU) : 7518
# Rec. Clause-clause subsumption calls : 2441
# Non-unit clause-clause subsumptions  : 83
# Unit Clause-clause subsumption calls : 305
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 14
# BW rewrite match successes           : 14
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 13169

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 28812 pages
