# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/040fbe255d6f4a018ca51f3e29b6d57c.E_input', d4_relat_1)).
fof(d11_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(v1_relat_1(X3)=>(X3=k7_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>(r2_hidden(X4,X2)&r2_hidden(k4_tarski(X4,X5),X1)))))), file('data/MPTP/proofs/040fbe255d6f4a018ca51f3e29b6d57c.E_input', d11_relat_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/040fbe255d6f4a018ca51f3e29b6d57c.E_input', d3_tarski)).
fof(dt_k7_relat_1, axiom, ![X1, X2]:(v1_relat_1(X1)=>v1_relat_1(k7_relat_1(X1,X2))), file('data/MPTP/proofs/040fbe255d6f4a018ca51f3e29b6d57c.E_input', dt_k7_relat_1)).
fof(t87_relat_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k1_relat_1(k7_relat_1(X2,X1)),X1)), file('data/MPTP/proofs/040fbe255d6f4a018ca51f3e29b6d57c.E_input', t87_relat_1)).
fof(c_0_5, plain, ![X10, X11, X12, X14, X15, X16, X17, X19]:(((~r2_hidden(X12,X11)|r2_hidden(k4_tarski(X12,esk3_3(X10,X11,X12)),X10)|X11!=k1_relat_1(X10))&(~r2_hidden(k4_tarski(X14,X15),X10)|r2_hidden(X14,X11)|X11!=k1_relat_1(X10)))&((~r2_hidden(esk4_2(X16,X17),X17)|~r2_hidden(k4_tarski(esk4_2(X16,X17),X19),X16)|X17=k1_relat_1(X16))&(r2_hidden(esk4_2(X16,X17),X17)|r2_hidden(k4_tarski(esk4_2(X16,X17),esk5_2(X16,X17)),X16)|X17=k1_relat_1(X16)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
fof(c_0_6, plain, ![X21, X22, X23, X24, X25, X26, X27]:((((r2_hidden(X24,X22)|~r2_hidden(k4_tarski(X24,X25),X23)|X23!=k7_relat_1(X21,X22)|~v1_relat_1(X23)|~v1_relat_1(X21))&(r2_hidden(k4_tarski(X24,X25),X21)|~r2_hidden(k4_tarski(X24,X25),X23)|X23!=k7_relat_1(X21,X22)|~v1_relat_1(X23)|~v1_relat_1(X21)))&(~r2_hidden(X26,X22)|~r2_hidden(k4_tarski(X26,X27),X21)|r2_hidden(k4_tarski(X26,X27),X23)|X23!=k7_relat_1(X21,X22)|~v1_relat_1(X23)|~v1_relat_1(X21)))&((~r2_hidden(k4_tarski(esk6_3(X21,X22,X23),esk7_3(X21,X22,X23)),X23)|(~r2_hidden(esk6_3(X21,X22,X23),X22)|~r2_hidden(k4_tarski(esk6_3(X21,X22,X23),esk7_3(X21,X22,X23)),X21))|X23=k7_relat_1(X21,X22)|~v1_relat_1(X23)|~v1_relat_1(X21))&((r2_hidden(esk6_3(X21,X22,X23),X22)|r2_hidden(k4_tarski(esk6_3(X21,X22,X23),esk7_3(X21,X22,X23)),X23)|X23=k7_relat_1(X21,X22)|~v1_relat_1(X23)|~v1_relat_1(X21))&(r2_hidden(k4_tarski(esk6_3(X21,X22,X23),esk7_3(X21,X22,X23)),X21)|r2_hidden(k4_tarski(esk6_3(X21,X22,X23),esk7_3(X21,X22,X23)),X23)|X23=k7_relat_1(X21,X22)|~v1_relat_1(X23)|~v1_relat_1(X21))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_relat_1])])])])])])).
cnf(c_0_7, plain, (r2_hidden(k4_tarski(X1,esk3_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),X4)|X4!=k7_relat_1(X5,X2)|~v1_relat_1(X4)|~v1_relat_1(X5)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, plain, (r2_hidden(k4_tarski(X1,esk3_3(X2,k1_relat_1(X2),X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_7])).
fof(c_0_10, plain, ![X30, X31, X32, X33, X34]:((~r1_tarski(X30,X31)|(~r2_hidden(X32,X30)|r2_hidden(X32,X31)))&((r2_hidden(esk8_2(X33,X34),X33)|r1_tarski(X33,X34))&(~r2_hidden(esk8_2(X33,X34),X34)|r1_tarski(X33,X34)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_11, plain, (r2_hidden(X1,X2)|X3!=k7_relat_1(X4,X2)|~r2_hidden(X1,k1_relat_1(X3))|~v1_relat_1(X4)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_12, plain, (r2_hidden(esk8_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_13, plain, ![X8, X9]:(~v1_relat_1(X8)|v1_relat_1(k7_relat_1(X8,X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_relat_1])])).
fof(c_0_14, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k1_relat_1(k7_relat_1(X2,X1)),X1))), inference(assume_negation,[status(cth)],[t87_relat_1])).
cnf(c_0_15, plain, (r2_hidden(esk8_2(k1_relat_1(X1),X2),X3)|r1_tarski(k1_relat_1(X1),X2)|X1!=k7_relat_1(X4,X3)|~v1_relat_1(X4)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16, plain, (v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_17, negated_conjecture, (v1_relat_1(esk2_0)&~r1_tarski(k1_relat_1(k7_relat_1(esk2_0,esk1_0)),esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
cnf(c_0_18, plain, (r2_hidden(esk8_2(k1_relat_1(k7_relat_1(X1,X2)),X3),X2)|r1_tarski(k1_relat_1(k7_relat_1(X1,X2)),X3)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_15]), c_0_16])).
cnf(c_0_19, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20, plain, (r1_tarski(X1,X2)|~r2_hidden(esk8_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_21, negated_conjecture, (r2_hidden(esk8_2(k1_relat_1(k7_relat_1(esk2_0,X1)),X2),X1)|r1_tarski(k1_relat_1(k7_relat_1(esk2_0,X1)),X2)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_22, negated_conjecture, (~r1_tarski(k1_relat_1(k7_relat_1(esk2_0,esk1_0)),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, negated_conjecture, (r1_tarski(k1_relat_1(k7_relat_1(esk2_0,X1)),X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_24, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_23])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 25
# Proof object clause steps            : 14
# Proof object formula steps           : 11
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 7
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 3
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 16
# Processed clauses                    : 44
# ...of these trivial                  : 1
# ...subsumed                          : 0
# ...remaining for further processing  : 43
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 48
# ...of the previous two non-trivial   : 43
# Contextual simplify-reflections      : 1
# Paramodulations                      : 43
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 26
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 23
# Current number of unprocessed clauses: 31
# ...number of literals in the above   : 146
# Current number of archived formulas  : 0
# Current number of archived clauses   : 17
# Clause-clause subsumption calls (NU) : 328
# Rec. Clause-clause subsumption calls : 72
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 24
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2350

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28636 pages
