# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t15_mcart_1, conjecture, ![X1, X2, X3, X4]:(r2_hidden(X1,k2_zfmisc_1(k2_tarski(X2,X3),X4))=>((k1_mcart_1(X1)=X2|k1_mcart_1(X1)=X3)&r2_hidden(k2_mcart_1(X1),X4))), file('data/MPTP/proofs/0412515f23d9461fbd007a7d9f8362c6.E_input', t15_mcart_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/0412515f23d9461fbd007a7d9f8362c6.E_input', d5_xboole_0)).
fof(t72_zfmisc_1, axiom, ![X1, X2, X3]:(k4_xboole_0(k2_tarski(X1,X2),X3)=k2_tarski(X1,X2)<=>(~(r2_hidden(X1,X3))&~(r2_hidden(X2,X3)))), file('data/MPTP/proofs/0412515f23d9461fbd007a7d9f8362c6.E_input', t72_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/0412515f23d9461fbd007a7d9f8362c6.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/0412515f23d9461fbd007a7d9f8362c6.E_input', t71_enumset1)).
fof(t12_mcart_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(k1_tarski(X2),X3))=>(k1_mcart_1(X1)=X2&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/0412515f23d9461fbd007a7d9f8362c6.E_input', t12_mcart_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/0412515f23d9461fbd007a7d9f8362c6.E_input', t69_enumset1)).
fof(t10_mcart_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(r2_hidden(k1_mcart_1(X1),X2)&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/0412515f23d9461fbd007a7d9f8362c6.E_input', t10_mcart_1)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/0412515f23d9461fbd007a7d9f8362c6.E_input', l54_zfmisc_1)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/0412515f23d9461fbd007a7d9f8362c6.E_input', t7_mcart_1)).
fof(t38_zfmisc_1, axiom, ![X1, X2, X3]:(r1_tarski(k2_tarski(X1,X2),X3)<=>(r2_hidden(X1,X3)&r2_hidden(X2,X3))), file('data/MPTP/proofs/0412515f23d9461fbd007a7d9f8362c6.E_input', t38_zfmisc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/0412515f23d9461fbd007a7d9f8362c6.E_input', d10_xboole_0)).
fof(c_0_12, negated_conjecture, ~(![X1, X2, X3, X4]:(r2_hidden(X1,k2_zfmisc_1(k2_tarski(X2,X3),X4))=>((k1_mcart_1(X1)=X2|k1_mcart_1(X1)=X3)&r2_hidden(k2_mcart_1(X1),X4)))), inference(assume_negation,[status(cth)],[t15_mcart_1])).
fof(c_0_13, plain, ![X28, X29, X30, X31, X32, X33, X34, X35]:((((r2_hidden(X31,X28)|~r2_hidden(X31,X30)|X30!=k4_xboole_0(X28,X29))&(~r2_hidden(X31,X29)|~r2_hidden(X31,X30)|X30!=k4_xboole_0(X28,X29)))&(~r2_hidden(X32,X28)|r2_hidden(X32,X29)|r2_hidden(X32,X30)|X30!=k4_xboole_0(X28,X29)))&((~r2_hidden(esk5_3(X33,X34,X35),X35)|(~r2_hidden(esk5_3(X33,X34,X35),X33)|r2_hidden(esk5_3(X33,X34,X35),X34))|X35=k4_xboole_0(X33,X34))&((r2_hidden(esk5_3(X33,X34,X35),X33)|r2_hidden(esk5_3(X33,X34,X35),X35)|X35=k4_xboole_0(X33,X34))&(~r2_hidden(esk5_3(X33,X34,X35),X34)|r2_hidden(esk5_3(X33,X34,X35),X35)|X35=k4_xboole_0(X33,X34))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_14, plain, ![X18, X19, X20]:(((~r2_hidden(X18,X20)|k4_xboole_0(k2_tarski(X18,X19),X20)!=k2_tarski(X18,X19))&(~r2_hidden(X19,X20)|k4_xboole_0(k2_tarski(X18,X19),X20)!=k2_tarski(X18,X19)))&(r2_hidden(X18,X20)|r2_hidden(X19,X20)|k4_xboole_0(k2_tarski(X18,X19),X20)=k2_tarski(X18,X19))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t72_zfmisc_1])])])])).
fof(c_0_15, plain, ![X39, X40]:k1_enumset1(X39,X39,X40)=k2_tarski(X39,X40), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_16, plain, ![X14, X15, X16]:k2_enumset1(X14,X14,X15,X16)=k1_enumset1(X14,X15,X16), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_17, negated_conjecture, (r2_hidden(esk1_0,k2_zfmisc_1(k2_tarski(esk2_0,esk3_0),esk4_0))&((k1_mcart_1(esk1_0)!=esk2_0|~r2_hidden(k2_mcart_1(esk1_0),esk4_0))&(k1_mcart_1(esk1_0)!=esk3_0|~r2_hidden(k2_mcart_1(esk1_0),esk4_0)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).
fof(c_0_18, plain, ![X41, X42, X43]:((k1_mcart_1(X41)=X42|~r2_hidden(X41,k2_zfmisc_1(k1_tarski(X42),X43)))&(r2_hidden(k2_mcart_1(X41),X43)|~r2_hidden(X41,k2_zfmisc_1(k1_tarski(X42),X43)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_mcart_1])])])).
fof(c_0_19, plain, ![X17]:k2_tarski(X17,X17)=k1_tarski(X17), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_20, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (r2_hidden(X1,X2)|r2_hidden(X3,X2)|k4_xboole_0(k2_tarski(X1,X3),X2)=k2_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_24, plain, ![X11, X12, X13]:((r2_hidden(k1_mcart_1(X11),X12)|~r2_hidden(X11,k2_zfmisc_1(X12,X13)))&(r2_hidden(k2_mcart_1(X11),X13)|~r2_hidden(X11,k2_zfmisc_1(X12,X13)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_mcart_1])])])).
cnf(c_0_25, negated_conjecture, (r2_hidden(esk1_0,k2_zfmisc_1(k2_tarski(esk2_0,esk3_0),esk4_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (k1_mcart_1(X1)=X2|~r2_hidden(X1,k2_zfmisc_1(k1_tarski(X2),X3))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_28, plain, ![X24, X25, X26, X27]:(((r2_hidden(X24,X26)|~r2_hidden(k4_tarski(X24,X25),k2_zfmisc_1(X26,X27)))&(r2_hidden(X25,X27)|~r2_hidden(k4_tarski(X24,X25),k2_zfmisc_1(X26,X27))))&(~r2_hidden(X24,X26)|~r2_hidden(X25,X27)|r2_hidden(k4_tarski(X24,X25),k2_zfmisc_1(X26,X27)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
fof(c_0_29, plain, ![X9, X10]:(k1_mcart_1(k4_tarski(X9,X10))=X9&k2_mcart_1(k4_tarski(X9,X10))=X10), inference(variable_rename,[status(thm)],[t7_mcart_1])).
cnf(c_0_30, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_31, plain, (k4_xboole_0(k2_enumset1(X1,X1,X1,X3),X2)=k2_enumset1(X1,X1,X1,X3)|r2_hidden(X3,X2)|r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_22]), c_0_22]), c_0_23]), c_0_23])).
cnf(c_0_32, plain, (r2_hidden(k1_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, negated_conjecture, (r2_hidden(esk1_0,k2_zfmisc_1(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_22]), c_0_23])).
cnf(c_0_34, plain, (k1_mcart_1(X1)=X2|~r2_hidden(X1,k2_zfmisc_1(k2_enumset1(X2,X2,X2,X2),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_27]), c_0_22]), c_0_23])).
cnf(c_0_35, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, plain, (r2_hidden(X1,X2)|r2_hidden(X3,X2)|~r2_hidden(X4,k2_enumset1(X1,X1,X1,X3))|~r2_hidden(X4,X2)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_38, negated_conjecture, (r2_hidden(k1_mcart_1(esk1_0),k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
fof(c_0_39, plain, ![X21, X22, X23]:(((r2_hidden(X21,X23)|~r1_tarski(k2_tarski(X21,X22),X23))&(r2_hidden(X22,X23)|~r1_tarski(k2_tarski(X21,X22),X23)))&(~r2_hidden(X21,X23)|~r2_hidden(X22,X23)|r1_tarski(k2_tarski(X21,X22),X23))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t38_zfmisc_1])])])).
fof(c_0_40, plain, ![X37, X38]:(((r1_tarski(X37,X38)|X37!=X38)&(r1_tarski(X38,X37)|X37!=X38))&(~r1_tarski(X37,X38)|~r1_tarski(X38,X37)|X37=X38)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_41, plain, (X1=X2|~r2_hidden(X1,k2_enumset1(X2,X2,X2,X2))|~r2_hidden(X3,X4)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])).
cnf(c_0_42, negated_conjecture, (r2_hidden(esk3_0,X1)|r2_hidden(esk2_0,X1)|~r2_hidden(k1_mcart_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_43, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_tarski(X3,X1),X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_44, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_45, plain, (r2_hidden(k2_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_46, negated_conjecture, (esk3_0=X1|r2_hidden(esk2_0,k2_enumset1(X1,X1,X1,X1))|~r2_hidden(k1_mcart_1(esk1_0),k2_enumset1(X1,X1,X1,X1))|~r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_47, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_enumset1(X3,X3,X3,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_22]), c_0_23])).
cnf(c_0_48, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_49, negated_conjecture, (k1_mcart_1(esk1_0)!=esk3_0|~r2_hidden(k2_mcart_1(esk1_0),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_50, negated_conjecture, (r2_hidden(k2_mcart_1(esk1_0),esk4_0)), inference(spm,[status(thm)],[c_0_45, c_0_33])).
cnf(c_0_51, negated_conjecture, (esk3_0=X1|r2_hidden(esk2_0,k2_enumset1(X1,X1,X1,X1))|~r2_hidden(k1_mcart_1(esk1_0),k2_enumset1(X1,X1,X1,X1))), inference(spm,[status(thm)],[c_0_46, c_0_33])).
cnf(c_0_52, plain, (r2_hidden(X1,k2_enumset1(X2,X2,X2,X1))), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_53, negated_conjecture, (esk3_0!=k1_mcart_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_50])])).
cnf(c_0_54, negated_conjecture, (k1_mcart_1(esk1_0)!=esk2_0|~r2_hidden(k2_mcart_1(esk1_0),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_55, negated_conjecture, (r2_hidden(esk2_0,k2_enumset1(k1_mcart_1(esk1_0),k1_mcart_1(esk1_0),k1_mcart_1(esk1_0),k1_mcart_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53])).
cnf(c_0_56, negated_conjecture, (esk2_0!=k1_mcart_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_50])])).
cnf(c_0_57, negated_conjecture, (~r2_hidden(X1,X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_55]), c_0_56])).
cnf(c_0_58, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_33, c_0_57]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 59
# Proof object clause steps            : 34
# Proof object formula steps           : 25
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 20
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 27
# Processed clauses                    : 91
# ...of these trivial                  : 1
# ...subsumed                          : 5
# ...remaining for further processing  : 85
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 3
# Generated clauses                    : 136
# ...of the previous two non-trivial   : 107
# Contextual simplify-reflections      : 0
# Paramodulations                      : 118
# Factorizations                       : 2
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 44
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 38
# Current number of unprocessed clauses: 56
# ...number of literals in the above   : 201
# Current number of archived formulas  : 0
# Current number of archived clauses   : 42
# Clause-clause subsumption calls (NU) : 291
# Rec. Clause-clause subsumption calls : 216
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 38
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3641

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28640 pages
