# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t66_zfmisc_1, conjecture, ![X1, X2]:~((k4_xboole_0(X1,k1_tarski(X2))=k1_xboole_0&(~(X1=k1_xboole_0)&~(X1=k1_tarski(X2))))), file('data/MPTP/proofs/0430efee19444fed815e2870d16d6ccf.E_input', t66_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/0430efee19444fed815e2870d16d6ccf.E_input', t69_enumset1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/0430efee19444fed815e2870d16d6ccf.E_input', d5_xboole_0)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/0430efee19444fed815e2870d16d6ccf.E_input', d1_tarski)).
fof(l35_zfmisc_1, axiom, ![X1, X2]:(k4_xboole_0(k1_tarski(X1),X2)=k1_xboole_0<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/0430efee19444fed815e2870d16d6ccf.E_input', l35_zfmisc_1)).
fof(l44_zfmisc_1, axiom, ![X1, X2]:~((~(X1=k1_tarski(X2))&(~(X1=k1_xboole_0)&![X3]:~((r2_hidden(X3,X1)&~(X3=X2)))))), file('data/MPTP/proofs/0430efee19444fed815e2870d16d6ccf.E_input', l44_zfmisc_1)).
fof(c_0_6, negated_conjecture, ~(![X1, X2]:~((k4_xboole_0(X1,k1_tarski(X2))=k1_xboole_0&(~(X1=k1_xboole_0)&~(X1=k1_tarski(X2)))))), inference(assume_negation,[status(cth)],[t66_zfmisc_1])).
fof(c_0_7, negated_conjecture, (k4_xboole_0(esk1_0,k1_tarski(esk2_0))=k1_xboole_0&(esk1_0!=k1_xboole_0&esk1_0!=k1_tarski(esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_6])])])])).
fof(c_0_8, plain, ![X7]:k2_tarski(X7,X7)=k1_tarski(X7), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_9, plain, ![X8, X9, X10, X11, X12, X13, X14, X15]:((((r2_hidden(X11,X8)|~r2_hidden(X11,X10)|X10!=k4_xboole_0(X8,X9))&(~r2_hidden(X11,X9)|~r2_hidden(X11,X10)|X10!=k4_xboole_0(X8,X9)))&(~r2_hidden(X12,X8)|r2_hidden(X12,X9)|r2_hidden(X12,X10)|X10!=k4_xboole_0(X8,X9)))&((~r2_hidden(esk3_3(X13,X14,X15),X15)|(~r2_hidden(esk3_3(X13,X14,X15),X13)|r2_hidden(esk3_3(X13,X14,X15),X14))|X15=k4_xboole_0(X13,X14))&((r2_hidden(esk3_3(X13,X14,X15),X13)|r2_hidden(esk3_3(X13,X14,X15),X15)|X15=k4_xboole_0(X13,X14))&(~r2_hidden(esk3_3(X13,X14,X15),X14)|r2_hidden(esk3_3(X13,X14,X15),X15)|X15=k4_xboole_0(X13,X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_10, plain, ![X17, X18, X19, X20, X21, X22]:(((~r2_hidden(X19,X18)|X19=X17|X18!=k1_tarski(X17))&(X20!=X17|r2_hidden(X20,X18)|X18!=k1_tarski(X17)))&((~r2_hidden(esk4_2(X21,X22),X22)|esk4_2(X21,X22)!=X21|X22=k1_tarski(X21))&(r2_hidden(esk4_2(X21,X22),X22)|esk4_2(X21,X22)=X21|X22=k1_tarski(X21)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_11, negated_conjecture, (k4_xboole_0(esk1_0,k1_tarski(esk2_0))=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (r2_hidden(esk3_3(X1,X2,X3),X1)|r2_hidden(esk3_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk2_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16, plain, (k4_xboole_0(X1,X2)=X1|r2_hidden(esk3_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_13])).
cnf(c_0_17, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_18, plain, ![X24, X25]:((k4_xboole_0(k1_tarski(X24),X25)!=k1_xboole_0|r2_hidden(X24,X25))&(~r2_hidden(X24,X25)|k4_xboole_0(k1_tarski(X24),X25)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l35_zfmisc_1])])).
cnf(c_0_19, plain, (X1=X3|X2!=k2_tarski(X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_14, c_0_12])).
cnf(c_0_20, plain, (r2_hidden(esk3_3(X1,X2,X3),X2)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk3_3(X1,X2,X3),X3)|~r2_hidden(esk3_3(X1,X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_21, negated_conjecture, (r2_hidden(esk3_3(esk1_0,k2_tarski(esk2_0,esk2_0),esk1_0),esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])).
cnf(c_0_22, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_23, plain, (k4_xboole_0(k1_tarski(X1),X2)=k1_xboole_0|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (X1=X2|~r2_hidden(X1,k2_tarski(X2,X2))), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_25, negated_conjecture, (r2_hidden(esk3_3(esk1_0,k2_tarski(esk2_0,esk2_0),esk1_0),k2_tarski(esk2_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_15]), c_0_21])]), c_0_17])).
cnf(c_0_26, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_27, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_28, plain, (k4_xboole_0(k2_tarski(X1,X1),X2)=k1_xboole_0|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_23, c_0_12])).
cnf(c_0_29, negated_conjecture, (esk3_3(esk1_0,k2_tarski(esk2_0,esk2_0),esk1_0)=esk2_0), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_30, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_26])).
cnf(c_0_31, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_32, plain, (r2_hidden(X1,k2_tarski(X2,X2))|~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_33, negated_conjecture, (r2_hidden(esk2_0,esk1_0)), inference(rw,[status(thm)],[c_0_21, c_0_29])).
cnf(c_0_34, negated_conjecture, (~r2_hidden(X1,k2_tarski(esk2_0,esk2_0))|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_30, c_0_15])).
fof(c_0_35, plain, ![X26, X27]:((r2_hidden(esk5_2(X26,X27),X26)|X26=k1_xboole_0|X26=k1_tarski(X27))&(esk5_2(X26,X27)!=X27|X26=k1_xboole_0|X26=k1_tarski(X27))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l44_zfmisc_1])])])])])).
cnf(c_0_36, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_31])).
cnf(c_0_37, negated_conjecture, (~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])).
cnf(c_0_38, negated_conjecture, (esk1_0!=k1_tarski(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_39, plain, (r2_hidden(esk5_2(X1,X2),X1)|X1=k1_xboole_0|X1=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_40, negated_conjecture, (r2_hidden(X1,k2_tarski(esk2_0,esk2_0))|~r2_hidden(X1,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_15]), c_0_37])).
cnf(c_0_41, negated_conjecture, (esk1_0!=k2_tarski(esk2_0,esk2_0)), inference(rw,[status(thm)],[c_0_38, c_0_12])).
cnf(c_0_42, plain, (X1=k1_xboole_0|X1=k2_tarski(X2,X2)|r2_hidden(esk5_2(X1,X2),X1)), inference(rw,[status(thm)],[c_0_39, c_0_12])).
cnf(c_0_43, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|esk5_2(X1,X2)!=X2), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_44, negated_conjecture, (X1=esk2_0|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_24, c_0_40])).
cnf(c_0_45, negated_conjecture, (X1=k1_xboole_0|r2_hidden(esk5_2(X1,esk2_0),X1)|X1!=esk1_0), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_46, plain, (X1=k1_xboole_0|X1=k2_tarski(X2,X2)|esk5_2(X1,X2)!=X2), inference(rw,[status(thm)],[c_0_43, c_0_12])).
cnf(c_0_47, negated_conjecture, (esk5_2(esk1_0,esk2_0)=esk2_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_17])).
cnf(c_0_48, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_41]), c_0_17]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 49
# Proof object clause steps            : 36
# Proof object formula steps           : 13
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 17
# Processed clauses                    : 156
# ...of these trivial                  : 6
# ...subsumed                          : 31
# ...remaining for further processing  : 119
# Other redundant clauses eliminated   : 49
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 1510
# ...of the previous two non-trivial   : 1225
# Contextual simplify-reflections      : 3
# Paramodulations                      : 1438
# Factorizations                       : 11
# Equation resolutions                 : 61
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 98
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 87
# Current number of unprocessed clauses: 1100
# ...number of literals in the above   : 5309
# Current number of archived formulas  : 0
# Current number of archived clauses   : 21
# Clause-clause subsumption calls (NU) : 619
# Rec. Clause-clause subsumption calls : 467
# Non-unit clause-clause subsumptions  : 11
# Unit Clause-clause subsumption calls : 54
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 20185

# -------------------------------------------------
# User time                : 0.021 s
# System time              : 0.003 s
# Total time               : 0.024 s
# Maximum resident set size: 28724 pages
