# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/04510142fbb64082909faa59061f5efa.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/04510142fbb64082909faa59061f5efa.E_input', t70_enumset1)).
fof(d1_ordinal1, axiom, ![X1]:k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1)), file('data/MPTP/proofs/04510142fbb64082909faa59061f5efa.E_input', d1_ordinal1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/04510142fbb64082909faa59061f5efa.E_input', t69_enumset1)).
fof(t12_ordinal1, conjecture, ![X1, X2]:(k1_ordinal1(X1)=k1_ordinal1(X2)=>X1=X2), file('data/MPTP/proofs/04510142fbb64082909faa59061f5efa.E_input', t12_ordinal1)).
fof(d3_xboole_0, axiom, ![X1, X2, X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2)))), file('data/MPTP/proofs/04510142fbb64082909faa59061f5efa.E_input', d3_xboole_0)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/04510142fbb64082909faa59061f5efa.E_input', t71_enumset1)).
fof(t10_ordinal1, axiom, ![X1]:r2_hidden(X1,k1_ordinal1(X1)), file('data/MPTP/proofs/04510142fbb64082909faa59061f5efa.E_input', t10_ordinal1)).
fof(d2_enumset1, axiom, ![X1, X2, X3, X4, X5]:(X5=k2_enumset1(X1,X2,X3,X4)<=>![X6]:(r2_hidden(X6,X5)<=>~((~(X6=X1)&(~(X6=X2)&(~(X6=X3)&~(X6=X4))))))), file('data/MPTP/proofs/04510142fbb64082909faa59061f5efa.E_input', d2_enumset1)).
fof(antisymmetry_r2_hidden, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>~(r2_hidden(X2,X1))), file('data/MPTP/proofs/04510142fbb64082909faa59061f5efa.E_input', antisymmetry_r2_hidden)).
fof(c_0_10, plain, ![X25, X26]:k3_tarski(k2_tarski(X25,X26))=k2_xboole_0(X25,X26), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_11, plain, ![X37, X38]:k1_enumset1(X37,X37,X38)=k2_tarski(X37,X38), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_12, plain, ![X40]:k1_ordinal1(X40)=k2_xboole_0(X40,k1_tarski(X40)), inference(variable_rename,[status(thm)],[d1_ordinal1])).
fof(c_0_13, plain, ![X39]:k2_tarski(X39,X39)=k1_tarski(X39), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_14, negated_conjecture, ~(![X1, X2]:(k1_ordinal1(X1)=k1_ordinal1(X2)=>X1=X2)), inference(assume_negation,[status(cth)],[t12_ordinal1])).
fof(c_0_15, plain, ![X28, X29, X30, X31, X32, X33, X34, X35]:(((~r2_hidden(X31,X30)|(r2_hidden(X31,X28)|r2_hidden(X31,X29))|X30!=k2_xboole_0(X28,X29))&((~r2_hidden(X32,X28)|r2_hidden(X32,X30)|X30!=k2_xboole_0(X28,X29))&(~r2_hidden(X32,X29)|r2_hidden(X32,X30)|X30!=k2_xboole_0(X28,X29))))&(((~r2_hidden(esk4_3(X33,X34,X35),X33)|~r2_hidden(esk4_3(X33,X34,X35),X35)|X35=k2_xboole_0(X33,X34))&(~r2_hidden(esk4_3(X33,X34,X35),X34)|~r2_hidden(esk4_3(X33,X34,X35),X35)|X35=k2_xboole_0(X33,X34)))&(r2_hidden(esk4_3(X33,X34,X35),X35)|(r2_hidden(esk4_3(X33,X34,X35),X33)|r2_hidden(esk4_3(X33,X34,X35),X34))|X35=k2_xboole_0(X33,X34)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])).
cnf(c_0_16, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_18, plain, ![X22, X23, X24]:k2_enumset1(X22,X22,X23,X24)=k1_enumset1(X22,X23,X24), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_19, plain, ![X27]:r2_hidden(X27,k1_ordinal1(X27)), inference(variable_rename,[status(thm)],[t10_ordinal1])).
cnf(c_0_20, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_22, negated_conjecture, (k1_ordinal1(esk1_0)=k1_ordinal1(esk2_0)&esk1_0!=esk2_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
cnf(c_0_23, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X2!=k2_xboole_0(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_25, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, plain, (r2_hidden(X1,k1_ordinal1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k2_tarski(X1,X1))), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_28, negated_conjecture, (k1_ordinal1(esk1_0)=k1_ordinal1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_29, plain, ![X9, X10, X11, X12, X13, X14, X15, X16, X17, X18, X19, X20]:(((~r2_hidden(X14,X13)|(X14=X9|(X14=X10|(X14=X11|X14=X12)))|X13!=k2_enumset1(X9,X10,X11,X12))&((X15!=X9|r2_hidden(X15,X13)|X13!=k2_enumset1(X9,X10,X11,X12))&((X15!=X10|r2_hidden(X15,X13)|X13!=k2_enumset1(X9,X10,X11,X12))&((X15!=X11|r2_hidden(X15,X13)|X13!=k2_enumset1(X9,X10,X11,X12))&(X15!=X12|r2_hidden(X15,X13)|X13!=k2_enumset1(X9,X10,X11,X12))))))&(((esk3_5(X16,X17,X18,X19,X20)!=X16|~r2_hidden(esk3_5(X16,X17,X18,X19,X20),X20)|X20=k2_enumset1(X16,X17,X18,X19))&((esk3_5(X16,X17,X18,X19,X20)!=X17|~r2_hidden(esk3_5(X16,X17,X18,X19,X20),X20)|X20=k2_enumset1(X16,X17,X18,X19))&((esk3_5(X16,X17,X18,X19,X20)!=X18|~r2_hidden(esk3_5(X16,X17,X18,X19,X20),X20)|X20=k2_enumset1(X16,X17,X18,X19))&(esk3_5(X16,X17,X18,X19,X20)!=X19|~r2_hidden(esk3_5(X16,X17,X18,X19,X20),X20)|X20=k2_enumset1(X16,X17,X18,X19)))))&(r2_hidden(esk3_5(X16,X17,X18,X19,X20),X20)|(esk3_5(X16,X17,X18,X19,X20)=X16|(esk3_5(X16,X17,X18,X19,X20)=X17|(esk3_5(X16,X17,X18,X19,X20)=X18|esk3_5(X16,X17,X18,X19,X20)=X19)))|X20=k2_enumset1(X16,X17,X18,X19)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_enumset1])])])])])])])).
cnf(c_0_30, plain, (r2_hidden(X1,X4)|r2_hidden(X1,X3)|X2!=k3_tarski(k2_enumset1(X3,X3,X3,X4))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
cnf(c_0_31, plain, (r2_hidden(X1,k3_tarski(k2_enumset1(X1,X1,X1,k2_enumset1(X1,X1,X1,X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_27]), c_0_17]), c_0_24]), c_0_25]), c_0_25])).
cnf(c_0_32, negated_conjecture, (k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)))=k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_27]), c_0_27]), c_0_17]), c_0_17]), c_0_24]), c_0_24]), c_0_25]), c_0_25]), c_0_25]), c_0_25])).
cnf(c_0_33, plain, (X1=X3|X1=X4|X1=X5|X1=X6|~r2_hidden(X1,X2)|X2!=k2_enumset1(X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k3_tarski(k2_enumset1(X3,X3,X3,X2)))), inference(er,[status(thm)],[c_0_30])).
cnf(c_0_35, negated_conjecture, (r2_hidden(esk1_0,k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
fof(c_0_36, plain, ![X41, X42]:(~r2_hidden(X41,X42)|~r2_hidden(X42,X41)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[antisymmetry_r2_hidden])])])).
cnf(c_0_37, plain, (X1=X2|X1=X3|X1=X4|X1=X5|~r2_hidden(X1,k2_enumset1(X5,X4,X3,X2))), inference(er,[status(thm)],[c_0_33])).
cnf(c_0_38, negated_conjecture, (r2_hidden(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))|r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_39, negated_conjecture, (esk1_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_40, plain, (~r2_hidden(X1,X2)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_41, negated_conjecture, (r2_hidden(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
cnf(c_0_42, negated_conjecture, (r2_hidden(X1,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))|r2_hidden(X1,esk1_0)|~r2_hidden(X1,k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))))), inference(spm,[status(thm)],[c_0_34, c_0_32])).
cnf(c_0_43, negated_conjecture, (~r2_hidden(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_44, negated_conjecture, (r2_hidden(esk2_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_31]), c_0_43])).
cnf(c_0_45, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_44]), c_0_39]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 46
# Proof object clause steps            : 25
# Proof object formula steps           : 21
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 24
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 20
# Processed clauses                    : 111
# ...of these trivial                  : 0
# ...subsumed                          : 25
# ...remaining for further processing  : 86
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 237
# ...of the previous two non-trivial   : 225
# Contextual simplify-reflections      : 0
# Paramodulations                      : 229
# Factorizations                       : 0
# Equation resolutions                 : 12
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 57
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 25
#    Non-unit-clauses                  : 23
# Current number of unprocessed clauses: 152
# ...number of literals in the above   : 269
# Current number of archived formulas  : 0
# Current number of archived clauses   : 26
# Clause-clause subsumption calls (NU) : 156
# Rec. Clause-clause subsumption calls : 148
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 146
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 13
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5118

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.003 s
# Total time               : 0.011 s
# Maximum resident set size: 28640 pages
