# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.007 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/04dbf87bd48d4af4972bc59a63fc2abf.E_input', d4_xboole_0)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/04dbf87bd48d4af4972bc59a63fc2abf.E_input', t12_setfam_1)).
fof(t124_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k8_relat_1(X1,X2)=k3_xboole_0(X2,k2_zfmisc_1(k1_relat_1(X2),X1))), file('data/MPTP/proofs/04dbf87bd48d4af4972bc59a63fc2abf.E_input', t124_relat_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/04dbf87bd48d4af4972bc59a63fc2abf.E_input', d10_xboole_0)).
fof(l139_zfmisc_1, axiom, ![X1, X2, X3]:~((r2_hidden(X1,k2_zfmisc_1(X2,X3))&![X4, X5]:~(k4_tarski(X4,X5)=X1))), file('data/MPTP/proofs/04dbf87bd48d4af4972bc59a63fc2abf.E_input', l139_zfmisc_1)).
fof(t125_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(k2_relat_1(X2),X1)=>k8_relat_1(X1,X2)=X2)), file('data/MPTP/proofs/04dbf87bd48d4af4972bc59a63fc2abf.E_input', t125_relat_1)).
fof(t23_mcart_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>(r2_hidden(X1,X2)=>X1=k4_tarski(k1_mcart_1(X1),k2_mcart_1(X1)))), file('data/MPTP/proofs/04dbf87bd48d4af4972bc59a63fc2abf.E_input', t23_mcart_1)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/04dbf87bd48d4af4972bc59a63fc2abf.E_input', t7_mcart_1)).
fof(c_0_8, plain, ![X17, X18, X19, X20, X21, X22, X23, X24]:((((r2_hidden(X20,X17)|~r2_hidden(X20,X19)|X19!=k3_xboole_0(X17,X18))&(r2_hidden(X20,X18)|~r2_hidden(X20,X19)|X19!=k3_xboole_0(X17,X18)))&(~r2_hidden(X21,X17)|~r2_hidden(X21,X18)|r2_hidden(X21,X19)|X19!=k3_xboole_0(X17,X18)))&((~r2_hidden(esk5_3(X22,X23,X24),X24)|(~r2_hidden(esk5_3(X22,X23,X24),X22)|~r2_hidden(esk5_3(X22,X23,X24),X23))|X24=k3_xboole_0(X22,X23))&((r2_hidden(esk5_3(X22,X23,X24),X22)|r2_hidden(esk5_3(X22,X23,X24),X24)|X24=k3_xboole_0(X22,X23))&(r2_hidden(esk5_3(X22,X23,X24),X23)|r2_hidden(esk5_3(X22,X23,X24),X24)|X24=k3_xboole_0(X22,X23))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_9, plain, ![X26, X27]:k1_setfam_1(k2_tarski(X26,X27))=k3_xboole_0(X26,X27), inference(variable_rename,[status(thm)],[t12_setfam_1])).
cnf(c_0_10, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_12, plain, ![X28, X29]:(~v1_relat_1(X29)|k8_relat_1(X28,X29)=k3_xboole_0(X29,k2_zfmisc_1(k1_relat_1(X29),X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t124_relat_1])])).
cnf(c_0_13, plain, (r2_hidden(X1,X2)|X3!=k1_setfam_1(k2_tarski(X4,X2))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_14, plain, (k8_relat_1(X2,X1)=k3_xboole_0(X1,k2_zfmisc_1(k1_relat_1(X1),X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_15, plain, ![X30, X31]:(((r1_tarski(X30,X31)|X30!=X31)&(r1_tarski(X31,X30)|X30!=X31))&(~r1_tarski(X30,X31)|~r1_tarski(X31,X30)|X30=X31)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_16, plain, ![X10, X11, X12]:(~r2_hidden(X10,k2_zfmisc_1(X11,X12))|k4_tarski(esk3_1(X10),esk4_1(X10))=X10), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l139_zfmisc_1])])])])])])).
cnf(c_0_17, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_setfam_1(k2_tarski(X3,X2)))), inference(er,[status(thm)],[c_0_13])).
cnf(c_0_18, plain, (k8_relat_1(X2,X1)=k1_setfam_1(k2_tarski(X1,k2_zfmisc_1(k1_relat_1(X1),X2)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_14, c_0_11])).
fof(c_0_19, plain, ![X8, X9]:(~v1_relat_1(X9)|(~r1_tarski(k2_relat_1(X9),X8)|k8_relat_1(X8,X9)=X9)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t125_relat_1])])).
cnf(c_0_20, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (k4_tarski(esk3_1(X1),esk4_1(X1))=X1|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (r2_hidden(X1,k2_zfmisc_1(k1_relat_1(X2),X3))|~r2_hidden(X1,k8_relat_1(X3,X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_23, plain, (k8_relat_1(X2,X1)=X1|~v1_relat_1(X1)|~r1_tarski(k2_relat_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_20])).
fof(c_0_25, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>(r2_hidden(X1,X2)=>X1=k4_tarski(k1_mcart_1(X1),k2_mcart_1(X1))))), inference(assume_negation,[status(cth)],[t23_mcart_1])).
cnf(c_0_26, plain, (k4_tarski(esk3_1(X1),esk4_1(X1))=X1|~r2_hidden(X1,k8_relat_1(X2,X3))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27, plain, (k8_relat_1(k2_relat_1(X1),X1)=X1|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_28, negated_conjecture, (v1_relat_1(esk2_0)&(r2_hidden(esk1_0,esk2_0)&esk1_0!=k4_tarski(k1_mcart_1(esk1_0),k2_mcart_1(esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])).
fof(c_0_29, plain, ![X15, X16]:(k1_mcart_1(k4_tarski(X15,X16))=X15&k2_mcart_1(k4_tarski(X15,X16))=X16), inference(variable_rename,[status(thm)],[t7_mcart_1])).
cnf(c_0_30, plain, (k4_tarski(esk3_1(X1),esk4_1(X1))=X1|~r2_hidden(X1,X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_31, negated_conjecture, (r2_hidden(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_32, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34, negated_conjecture, (k4_tarski(esk3_1(esk1_0),esk4_1(esk1_0))=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])])).
cnf(c_0_35, negated_conjecture, (esk3_1(esk1_0)=k1_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_36, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, negated_conjecture, (k4_tarski(k1_mcart_1(esk1_0),esk4_1(esk1_0))=esk1_0), inference(rw,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_38, negated_conjecture, (esk4_1(esk1_0)=k2_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_39, negated_conjecture, (esk1_0!=k4_tarski(k1_mcart_1(esk1_0),k2_mcart_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_40, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_38]), c_0_39]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 41
# Proof object clause steps            : 24
# Proof object formula steps           : 17
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 9
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 17
# Processed clauses                    : 46
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 46
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 52
# ...of the previous two non-trivial   : 47
# Contextual simplify-reflections      : 0
# Paramodulations                      : 43
# Factorizations                       : 4
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 23
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 15
# Current number of unprocessed clauses: 33
# ...number of literals in the above   : 114
# Current number of archived formulas  : 0
# Current number of archived clauses   : 19
# Clause-clause subsumption calls (NU) : 39
# Rec. Clause-clause subsumption calls : 31
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 18
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1807

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.000 s
# Total time               : 0.008 s
# Maximum resident set size: 28644 pages
