# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(redefinition_r2_relset_1, axiom, ![X1, X2, X3, X4]:((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>(r2_relset_1(X1,X2,X3,X4)<=>X3=X4)), file('data/MPTP/proofs/0500efce91ca48719932e9ab751352a4.E_input', redefinition_r2_relset_1)).
fof(dt_k1_toler_1, axiom, ![X1, X2]:(v1_relat_1(X1)=>m1_subset_1(k1_toler_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X2,X2)))), file('data/MPTP/proofs/0500efce91ca48719932e9ab751352a4.E_input', dt_k1_toler_1)).
fof(d7_waybel_9, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>![X4]:((v6_waybel_0(X4,X1)&l1_waybel_0(X4,X1))=>(X4=k4_waybel_9(X1,X2,X3)<=>(![X5]:(r2_hidden(X5,u1_struct_0(X4))<=>?[X6]:(m1_subset_1(X6,u1_struct_0(X2))&(X6=X5&r1_orders_2(X2,X3,X6))))&(r2_relset_1(u1_struct_0(X4),u1_struct_0(X4),u1_orders_2(X4),k1_toler_1(u1_orders_2(X2),u1_struct_0(X4)))&u1_waybel_0(X1,X4)=k2_partfun1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2),u1_struct_0(X4))))))))), file('data/MPTP/proofs/0500efce91ca48719932e9ab751352a4.E_input', d7_waybel_9)).
fof(dt_k4_waybel_9, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))&m1_subset_1(X3,u1_struct_0(X2))))=>(v6_waybel_0(k4_waybel_9(X1,X2,X3),X1)&l1_waybel_0(k4_waybel_9(X1,X2,X3),X1))), file('data/MPTP/proofs/0500efce91ca48719932e9ab751352a4.E_input', dt_k4_waybel_9)).
fof(dt_l1_waybel_0, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(l1_waybel_0(X2,X1)=>l1_orders_2(X2))), file('data/MPTP/proofs/0500efce91ca48719932e9ab751352a4.E_input', dt_l1_waybel_0)).
fof(t14_waybel_9, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>(v2_yellow_6(k4_waybel_9(X1,X2,X3),X1,X2)&m1_yellow_6(k4_waybel_9(X1,X2,X3),X1,X2))))), file('data/MPTP/proofs/0500efce91ca48719932e9ab751352a4.E_input', t14_waybel_9)).
fof(dt_u1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))), file('data/MPTP/proofs/0500efce91ca48719932e9ab751352a4.E_input', dt_u1_orders_2)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/0500efce91ca48719932e9ab751352a4.E_input', cc1_relset_1)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/0500efce91ca48719932e9ab751352a4.E_input', t17_xboole_1)).
fof(d6_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:k2_wellord1(X1,X2)=k3_xboole_0(X1,k2_zfmisc_1(X2,X2))), file('data/MPTP/proofs/0500efce91ca48719932e9ab751352a4.E_input', d6_wellord1)).
fof(d9_yellow_6, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(l1_waybel_0(X2,X1)=>![X3]:(m1_yellow_6(X3,X1,X2)=>(v2_yellow_6(X3,X1,X2)<=>(v4_yellow_0(X3,X2)&m1_yellow_0(X3,X2)))))), file('data/MPTP/proofs/0500efce91ca48719932e9ab751352a4.E_input', d9_yellow_6)).
fof(d14_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_yellow_0(X2,X1)=>(v4_yellow_0(X2,X1)<=>u1_orders_2(X2)=k1_toler_1(u1_orders_2(X1),u1_struct_0(X2))))), file('data/MPTP/proofs/0500efce91ca48719932e9ab751352a4.E_input', d14_yellow_0)).
fof(redefinition_k1_toler_1, axiom, ![X1, X2]:(v1_relat_1(X1)=>k1_toler_1(X1,X2)=k2_wellord1(X1,X2)), file('data/MPTP/proofs/0500efce91ca48719932e9ab751352a4.E_input', redefinition_k1_toler_1)).
fof(d8_yellow_6, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(l1_waybel_0(X2,X1)=>![X3]:(l1_waybel_0(X3,X1)=>(m1_yellow_6(X3,X1,X2)<=>(m1_yellow_0(X3,X2)&u1_waybel_0(X1,X3)=k2_partfun1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2),u1_struct_0(X3))))))), file('data/MPTP/proofs/0500efce91ca48719932e9ab751352a4.E_input', d8_yellow_6)).
fof(d13_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(m1_yellow_0(X2,X1)<=>(r1_tarski(u1_struct_0(X2),u1_struct_0(X1))&r1_tarski(u1_orders_2(X2),u1_orders_2(X1)))))), file('data/MPTP/proofs/0500efce91ca48719932e9ab751352a4.E_input', d13_yellow_0)).
fof(t13_waybel_9, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>r1_tarski(u1_struct_0(k4_waybel_9(X1,X2,X3)),u1_struct_0(X2))))), file('data/MPTP/proofs/0500efce91ca48719932e9ab751352a4.E_input', t13_waybel_9)).
fof(c_0_16, plain, ![X7, X8, X9, X10]:((~r2_relset_1(X7,X8,X9,X10)|X9=X10|(~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))&(X9!=X10|r2_relset_1(X7,X8,X9,X10)|(~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r2_relset_1])])])).
fof(c_0_17, plain, ![X11, X12]:(~v1_relat_1(X11)|m1_subset_1(k1_toler_1(X11,X12),k1_zfmisc_1(k2_zfmisc_1(X12,X12)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_toler_1])])).
fof(c_0_18, plain, ![X13, X14, X15, X16, X17, X19, X20, X22]:(((((m1_subset_1(esk1_5(X13,X14,X15,X16,X17),u1_struct_0(X14))|~r2_hidden(X17,u1_struct_0(X16))|X16!=k4_waybel_9(X13,X14,X15)|(~v6_waybel_0(X16,X13)|~l1_waybel_0(X16,X13))|~m1_subset_1(X15,u1_struct_0(X14))|(v2_struct_0(X14)|~l1_waybel_0(X14,X13))|(v2_struct_0(X13)|~l1_struct_0(X13)))&((esk1_5(X13,X14,X15,X16,X17)=X17|~r2_hidden(X17,u1_struct_0(X16))|X16!=k4_waybel_9(X13,X14,X15)|(~v6_waybel_0(X16,X13)|~l1_waybel_0(X16,X13))|~m1_subset_1(X15,u1_struct_0(X14))|(v2_struct_0(X14)|~l1_waybel_0(X14,X13))|(v2_struct_0(X13)|~l1_struct_0(X13)))&(r1_orders_2(X14,X15,esk1_5(X13,X14,X15,X16,X17))|~r2_hidden(X17,u1_struct_0(X16))|X16!=k4_waybel_9(X13,X14,X15)|(~v6_waybel_0(X16,X13)|~l1_waybel_0(X16,X13))|~m1_subset_1(X15,u1_struct_0(X14))|(v2_struct_0(X14)|~l1_waybel_0(X14,X13))|(v2_struct_0(X13)|~l1_struct_0(X13)))))&(~m1_subset_1(X20,u1_struct_0(X14))|(X20!=X19|~r1_orders_2(X14,X15,X20))|r2_hidden(X19,u1_struct_0(X16))|X16!=k4_waybel_9(X13,X14,X15)|(~v6_waybel_0(X16,X13)|~l1_waybel_0(X16,X13))|~m1_subset_1(X15,u1_struct_0(X14))|(v2_struct_0(X14)|~l1_waybel_0(X14,X13))|(v2_struct_0(X13)|~l1_struct_0(X13))))&((r2_relset_1(u1_struct_0(X16),u1_struct_0(X16),u1_orders_2(X16),k1_toler_1(u1_orders_2(X14),u1_struct_0(X16)))|X16!=k4_waybel_9(X13,X14,X15)|(~v6_waybel_0(X16,X13)|~l1_waybel_0(X16,X13))|~m1_subset_1(X15,u1_struct_0(X14))|(v2_struct_0(X14)|~l1_waybel_0(X14,X13))|(v2_struct_0(X13)|~l1_struct_0(X13)))&(u1_waybel_0(X13,X16)=k2_partfun1(u1_struct_0(X14),u1_struct_0(X13),u1_waybel_0(X13,X14),u1_struct_0(X16))|X16!=k4_waybel_9(X13,X14,X15)|(~v6_waybel_0(X16,X13)|~l1_waybel_0(X16,X13))|~m1_subset_1(X15,u1_struct_0(X14))|(v2_struct_0(X14)|~l1_waybel_0(X14,X13))|(v2_struct_0(X13)|~l1_struct_0(X13)))))&((~r2_hidden(esk2_4(X13,X14,X15,X16),u1_struct_0(X16))|(~m1_subset_1(X22,u1_struct_0(X14))|(X22!=esk2_4(X13,X14,X15,X16)|~r1_orders_2(X14,X15,X22)))|(~r2_relset_1(u1_struct_0(X16),u1_struct_0(X16),u1_orders_2(X16),k1_toler_1(u1_orders_2(X14),u1_struct_0(X16)))|u1_waybel_0(X13,X16)!=k2_partfun1(u1_struct_0(X14),u1_struct_0(X13),u1_waybel_0(X13,X14),u1_struct_0(X16)))|X16=k4_waybel_9(X13,X14,X15)|(~v6_waybel_0(X16,X13)|~l1_waybel_0(X16,X13))|~m1_subset_1(X15,u1_struct_0(X14))|(v2_struct_0(X14)|~l1_waybel_0(X14,X13))|(v2_struct_0(X13)|~l1_struct_0(X13)))&((m1_subset_1(esk3_4(X13,X14,X15,X16),u1_struct_0(X14))|r2_hidden(esk2_4(X13,X14,X15,X16),u1_struct_0(X16))|(~r2_relset_1(u1_struct_0(X16),u1_struct_0(X16),u1_orders_2(X16),k1_toler_1(u1_orders_2(X14),u1_struct_0(X16)))|u1_waybel_0(X13,X16)!=k2_partfun1(u1_struct_0(X14),u1_struct_0(X13),u1_waybel_0(X13,X14),u1_struct_0(X16)))|X16=k4_waybel_9(X13,X14,X15)|(~v6_waybel_0(X16,X13)|~l1_waybel_0(X16,X13))|~m1_subset_1(X15,u1_struct_0(X14))|(v2_struct_0(X14)|~l1_waybel_0(X14,X13))|(v2_struct_0(X13)|~l1_struct_0(X13)))&((esk3_4(X13,X14,X15,X16)=esk2_4(X13,X14,X15,X16)|r2_hidden(esk2_4(X13,X14,X15,X16),u1_struct_0(X16))|(~r2_relset_1(u1_struct_0(X16),u1_struct_0(X16),u1_orders_2(X16),k1_toler_1(u1_orders_2(X14),u1_struct_0(X16)))|u1_waybel_0(X13,X16)!=k2_partfun1(u1_struct_0(X14),u1_struct_0(X13),u1_waybel_0(X13,X14),u1_struct_0(X16)))|X16=k4_waybel_9(X13,X14,X15)|(~v6_waybel_0(X16,X13)|~l1_waybel_0(X16,X13))|~m1_subset_1(X15,u1_struct_0(X14))|(v2_struct_0(X14)|~l1_waybel_0(X14,X13))|(v2_struct_0(X13)|~l1_struct_0(X13)))&(r1_orders_2(X14,X15,esk3_4(X13,X14,X15,X16))|r2_hidden(esk2_4(X13,X14,X15,X16),u1_struct_0(X16))|(~r2_relset_1(u1_struct_0(X16),u1_struct_0(X16),u1_orders_2(X16),k1_toler_1(u1_orders_2(X14),u1_struct_0(X16)))|u1_waybel_0(X13,X16)!=k2_partfun1(u1_struct_0(X14),u1_struct_0(X13),u1_waybel_0(X13,X14),u1_struct_0(X16)))|X16=k4_waybel_9(X13,X14,X15)|(~v6_waybel_0(X16,X13)|~l1_waybel_0(X16,X13))|~m1_subset_1(X15,u1_struct_0(X14))|(v2_struct_0(X14)|~l1_waybel_0(X14,X13))|(v2_struct_0(X13)|~l1_struct_0(X13))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_waybel_9])])])])])])])).
fof(c_0_19, plain, ![X24, X25, X26]:((v6_waybel_0(k4_waybel_9(X24,X25,X26),X24)|(v2_struct_0(X24)|~l1_struct_0(X24)|(v2_struct_0(X25)|~l1_waybel_0(X25,X24)|~m1_subset_1(X26,u1_struct_0(X25)))))&(l1_waybel_0(k4_waybel_9(X24,X25,X26),X24)|(v2_struct_0(X24)|~l1_struct_0(X24)|(v2_struct_0(X25)|~l1_waybel_0(X25,X24)|~m1_subset_1(X26,u1_struct_0(X25)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_waybel_9])])])])).
cnf(c_0_20, plain, (X3=X4|~r2_relset_1(X1,X2,X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, plain, (m1_subset_1(k1_toler_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, plain, (r2_relset_1(u1_struct_0(X1),u1_struct_0(X1),u1_orders_2(X1),k1_toler_1(u1_orders_2(X2),u1_struct_0(X1)))|v2_struct_0(X2)|v2_struct_0(X3)|X1!=k4_waybel_9(X3,X2,X4)|~v6_waybel_0(X1,X3)|~l1_waybel_0(X1,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_waybel_0(X2,X3)|~l1_struct_0(X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, plain, (l1_waybel_0(k4_waybel_9(X1,X2,X3),X1)|v2_struct_0(X1)|v2_struct_0(X2)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, plain, (v6_waybel_0(k4_waybel_9(X1,X2,X3),X1)|v2_struct_0(X1)|v2_struct_0(X2)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_25, plain, ![X27, X28]:(~l1_struct_0(X27)|(~l1_waybel_0(X28,X27)|l1_orders_2(X28))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_waybel_0])])])).
fof(c_0_26, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>(v2_yellow_6(k4_waybel_9(X1,X2,X3),X1,X2)&m1_yellow_6(k4_waybel_9(X1,X2,X3),X1,X2)))))), inference(assume_negation,[status(cth)],[t14_waybel_9])).
cnf(c_0_27, plain, (X1=k1_toler_1(X2,X3)|~v1_relat_1(X2)|~r2_relset_1(X3,X3,X1,k1_toler_1(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X3)))), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_28, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r2_relset_1(u1_struct_0(k4_waybel_9(X1,X2,X3)),u1_struct_0(k4_waybel_9(X1,X2,X3)),u1_orders_2(k4_waybel_9(X1,X2,X3)),k1_toler_1(u1_orders_2(X2),u1_struct_0(k4_waybel_9(X1,X2,X3))))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_22]), c_0_23]), c_0_24])).
fof(c_0_29, plain, ![X32]:(~l1_orders_2(X32)|m1_subset_1(u1_orders_2(X32),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X32),u1_struct_0(X32))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_orders_2])])).
cnf(c_0_30, plain, (l1_orders_2(X2)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_31, plain, ![X47, X48, X49]:(~m1_subset_1(X49,k1_zfmisc_1(k2_zfmisc_1(X47,X48)))|v1_relat_1(X49)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_32, plain, ![X33, X34]:r1_tarski(k3_xboole_0(X33,X34),X33), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_33, plain, ![X35, X36]:(~v1_relat_1(X35)|k2_wellord1(X35,X36)=k3_xboole_0(X35,k2_zfmisc_1(X36,X36))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_wellord1])])])).
fof(c_0_34, negated_conjecture, ((~v2_struct_0(esk4_0)&l1_struct_0(esk4_0))&((~v2_struct_0(esk5_0)&l1_waybel_0(esk5_0,esk4_0))&(m1_subset_1(esk6_0,u1_struct_0(esk5_0))&(~v2_yellow_6(k4_waybel_9(esk4_0,esk5_0,esk6_0),esk4_0,esk5_0)|~m1_yellow_6(k4_waybel_9(esk4_0,esk5_0,esk6_0),esk4_0,esk5_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_26])])])])).
fof(c_0_35, plain, ![X37, X38, X39]:(((v4_yellow_0(X39,X38)|~v2_yellow_6(X39,X37,X38)|~m1_yellow_6(X39,X37,X38)|~l1_waybel_0(X38,X37)|~l1_struct_0(X37))&(m1_yellow_0(X39,X38)|~v2_yellow_6(X39,X37,X38)|~m1_yellow_6(X39,X37,X38)|~l1_waybel_0(X38,X37)|~l1_struct_0(X37)))&(~v4_yellow_0(X39,X38)|~m1_yellow_0(X39,X38)|v2_yellow_6(X39,X37,X38)|~m1_yellow_6(X39,X37,X38)|~l1_waybel_0(X38,X37)|~l1_struct_0(X37))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_6])])])])).
fof(c_0_36, plain, ![X40, X41]:((~v4_yellow_0(X41,X40)|u1_orders_2(X41)=k1_toler_1(u1_orders_2(X40),u1_struct_0(X41))|~m1_yellow_0(X41,X40)|~l1_orders_2(X40))&(u1_orders_2(X41)!=k1_toler_1(u1_orders_2(X40),u1_struct_0(X41))|v4_yellow_0(X41,X40)|~m1_yellow_0(X41,X40)|~l1_orders_2(X40))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d14_yellow_0])])])])).
cnf(c_0_37, plain, (k1_toler_1(u1_orders_2(X1),u1_struct_0(k4_waybel_9(X2,X1,X3)))=u1_orders_2(k4_waybel_9(X2,X1,X3))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)|~v1_relat_1(u1_orders_2(X1))|~m1_subset_1(u1_orders_2(k4_waybel_9(X2,X1,X3)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k4_waybel_9(X2,X1,X3)),u1_struct_0(k4_waybel_9(X2,X1,X3)))))|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_38, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, plain, (l1_orders_2(k4_waybel_9(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(spm,[status(thm)],[c_0_30, c_0_23])).
cnf(c_0_40, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_42, plain, (k2_wellord1(X1,X2)=k3_xboole_0(X1,k2_zfmisc_1(X2,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_43, plain, ![X42, X43]:(~v1_relat_1(X42)|k1_toler_1(X42,X43)=k2_wellord1(X42,X43)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_toler_1])])).
cnf(c_0_44, negated_conjecture, (~v2_yellow_6(k4_waybel_9(esk4_0,esk5_0,esk6_0),esk4_0,esk5_0)|~m1_yellow_6(k4_waybel_9(esk4_0,esk5_0,esk6_0),esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45, plain, (v2_yellow_6(X1,X3,X2)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~m1_yellow_6(X1,X3,X2)|~l1_waybel_0(X2,X3)|~l1_struct_0(X3)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_46, negated_conjecture, (l1_waybel_0(esk5_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_47, negated_conjecture, (l1_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_48, plain, (v4_yellow_0(X1,X2)|u1_orders_2(X1)!=k1_toler_1(u1_orders_2(X2),u1_struct_0(X1))|~m1_yellow_0(X1,X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_49, plain, (k1_toler_1(u1_orders_2(X1),u1_struct_0(k4_waybel_9(X2,X1,X3)))=u1_orders_2(k4_waybel_9(X2,X1,X3))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)|~v1_relat_1(u1_orders_2(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
cnf(c_0_50, plain, (v1_relat_1(u1_orders_2(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_40, c_0_38])).
fof(c_0_51, plain, ![X44, X45, X46]:(((m1_yellow_0(X46,X45)|~m1_yellow_6(X46,X44,X45)|~l1_waybel_0(X46,X44)|~l1_waybel_0(X45,X44)|~l1_struct_0(X44))&(u1_waybel_0(X44,X46)=k2_partfun1(u1_struct_0(X45),u1_struct_0(X44),u1_waybel_0(X44,X45),u1_struct_0(X46))|~m1_yellow_6(X46,X44,X45)|~l1_waybel_0(X46,X44)|~l1_waybel_0(X45,X44)|~l1_struct_0(X44)))&(~m1_yellow_0(X46,X45)|u1_waybel_0(X44,X46)!=k2_partfun1(u1_struct_0(X45),u1_struct_0(X44),u1_waybel_0(X44,X45),u1_struct_0(X46))|m1_yellow_6(X46,X44,X45)|~l1_waybel_0(X46,X44)|~l1_waybel_0(X45,X44)|~l1_struct_0(X44))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_yellow_6])])])])).
cnf(c_0_52, plain, (u1_waybel_0(X1,X2)=k2_partfun1(u1_struct_0(X3),u1_struct_0(X1),u1_waybel_0(X1,X3),u1_struct_0(X2))|v2_struct_0(X3)|v2_struct_0(X1)|X2!=k4_waybel_9(X1,X3,X4)|~v6_waybel_0(X2,X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X4,u1_struct_0(X3))|~l1_waybel_0(X3,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_53, plain, (r1_tarski(k2_wellord1(X1,X2),X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_54, plain, (k1_toler_1(X1,X2)=k2_wellord1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_55, negated_conjecture, (~m1_yellow_0(k4_waybel_9(esk4_0,esk5_0,esk6_0),esk5_0)|~v4_yellow_0(k4_waybel_9(esk4_0,esk5_0,esk6_0),esk5_0)|~m1_yellow_6(k4_waybel_9(esk4_0,esk5_0,esk6_0),esk4_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46]), c_0_47])])).
cnf(c_0_56, plain, (v4_yellow_0(k4_waybel_9(X1,X2,X3),X2)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_yellow_0(k4_waybel_9(X1,X2,X3),X2)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50]), c_0_30])).
cnf(c_0_57, negated_conjecture, (m1_subset_1(esk6_0,u1_struct_0(esk5_0))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_58, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_59, negated_conjecture, (~v2_struct_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_60, plain, (m1_yellow_6(X1,X3,X2)|~m1_yellow_0(X1,X2)|u1_waybel_0(X3,X1)!=k2_partfun1(u1_struct_0(X2),u1_struct_0(X3),u1_waybel_0(X3,X2),u1_struct_0(X1))|~l1_waybel_0(X1,X3)|~l1_waybel_0(X2,X3)|~l1_struct_0(X3)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_61, plain, (k2_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_waybel_0(X2,X1),u1_struct_0(k4_waybel_9(X2,X1,X3)))=u1_waybel_0(X2,k4_waybel_9(X2,X1,X3))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_52]), c_0_23]), c_0_24])).
fof(c_0_62, plain, ![X50, X51]:(((r1_tarski(u1_struct_0(X51),u1_struct_0(X50))|~m1_yellow_0(X51,X50)|~l1_orders_2(X51)|~l1_orders_2(X50))&(r1_tarski(u1_orders_2(X51),u1_orders_2(X50))|~m1_yellow_0(X51,X50)|~l1_orders_2(X51)|~l1_orders_2(X50)))&(~r1_tarski(u1_struct_0(X51),u1_struct_0(X50))|~r1_tarski(u1_orders_2(X51),u1_orders_2(X50))|m1_yellow_0(X51,X50)|~l1_orders_2(X51)|~l1_orders_2(X50))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d13_yellow_0])])])])).
cnf(c_0_63, plain, (r1_tarski(k1_toler_1(X1,X2),X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
fof(c_0_64, plain, ![X52, X53, X54]:(v2_struct_0(X52)|~l1_struct_0(X52)|(v2_struct_0(X53)|~l1_waybel_0(X53,X52)|(~m1_subset_1(X54,u1_struct_0(X53))|r1_tarski(u1_struct_0(k4_waybel_9(X52,X53,X54)),u1_struct_0(X53))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t13_waybel_9])])])])).
cnf(c_0_65, negated_conjecture, (~m1_yellow_0(k4_waybel_9(esk4_0,esk5_0,esk6_0),esk5_0)|~m1_yellow_6(k4_waybel_9(esk4_0,esk5_0,esk6_0),esk4_0,esk5_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_46]), c_0_47]), c_0_57])]), c_0_58]), c_0_59])).
cnf(c_0_66, plain, (m1_yellow_6(k4_waybel_9(X1,X2,X3),X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_yellow_0(k4_waybel_9(X1,X2,X3),X2)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_23])).
cnf(c_0_67, plain, (m1_yellow_0(X1,X2)|~r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~r1_tarski(u1_orders_2(X1),u1_orders_2(X2))|~l1_orders_2(X1)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_68, plain, (r1_tarski(u1_orders_2(k4_waybel_9(X1,X2,X3)),u1_orders_2(X2))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)|~v1_relat_1(u1_orders_2(X2))|~m1_subset_1(X3,u1_struct_0(X2))), inference(spm,[status(thm)],[c_0_63, c_0_49])).
cnf(c_0_69, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(u1_struct_0(k4_waybel_9(X1,X2,X3)),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_70, negated_conjecture, (~m1_yellow_0(k4_waybel_9(esk4_0,esk5_0,esk6_0),esk5_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_46]), c_0_47]), c_0_57])]), c_0_59]), c_0_58])).
cnf(c_0_71, plain, (m1_yellow_0(k4_waybel_9(X1,X2,X3),X2)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_50]), c_0_30]), c_0_39]), c_0_69])).
cnf(c_0_72, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_46]), c_0_47]), c_0_57])]), c_0_59]), c_0_58]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 73
# Proof object clause steps            : 40
# Proof object formula steps           : 33
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 35
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 39
# Processed clauses                    : 118
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 117
# Other redundant clauses eliminated   : 9
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 0
# Generated clauses                    : 71
# ...of the previous two non-trivial   : 60
# Contextual simplify-reflections      : 28
# Paramodulations                      : 63
# Factorizations                       : 0
# Equation resolutions                 : 9
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 68
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 60
# Current number of unprocessed clauses: 20
# ...number of literals in the above   : 258
# Current number of archived formulas  : 0
# Current number of archived clauses   : 41
# Clause-clause subsumption calls (NU) : 2355
# Rec. Clause-clause subsumption calls : 165
# Non-unit clause-clause subsumptions  : 31
# Unit Clause-clause subsumption calls : 6
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7502

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.000 s
# Total time               : 0.015 s
# Maximum resident set size: 28680 pages
