# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/0507cbfdc01b45e59e5ee29c6e40d839.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/0507cbfdc01b45e59e5ee29c6e40d839.E_input', t70_enumset1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/0507cbfdc01b45e59e5ee29c6e40d839.E_input', d10_xboole_0)).
fof(t43_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3)), file('data/MPTP/proofs/0507cbfdc01b45e59e5ee29c6e40d839.E_input', t43_xboole_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/0507cbfdc01b45e59e5ee29c6e40d839.E_input', t7_xboole_1)).
fof(t143_zfmisc_1, conjecture, ![X1, X2, X3, X4, X5, X6]:((r1_tarski(X1,k2_zfmisc_1(X3,X4))&r1_tarski(X2,k2_zfmisc_1(X5,X6)))=>r1_tarski(k2_xboole_0(X1,X2),k2_zfmisc_1(k2_xboole_0(X3,X5),k2_xboole_0(X4,X6)))), file('data/MPTP/proofs/0507cbfdc01b45e59e5ee29c6e40d839.E_input', t143_zfmisc_1)).
fof(t106_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k4_xboole_0(X2,X3))=>(r1_tarski(X1,X2)&r1_xboole_0(X1,X3))), file('data/MPTP/proofs/0507cbfdc01b45e59e5ee29c6e40d839.E_input', t106_xboole_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/0507cbfdc01b45e59e5ee29c6e40d839.E_input', t2_xboole_1)).
fof(t63_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_xboole_0(X2,X3))=>r1_xboole_0(X1,X3)), file('data/MPTP/proofs/0507cbfdc01b45e59e5ee29c6e40d839.E_input', t63_xboole_1)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/0507cbfdc01b45e59e5ee29c6e40d839.E_input', symmetry_r1_xboole_0)).
fof(t86_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_xboole_0(X1,X3))=>r1_tarski(X1,k4_xboole_0(X2,X3))), file('data/MPTP/proofs/0507cbfdc01b45e59e5ee29c6e40d839.E_input', t86_xboole_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/0507cbfdc01b45e59e5ee29c6e40d839.E_input', t36_xboole_1)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/0507cbfdc01b45e59e5ee29c6e40d839.E_input', t98_xboole_1)).
fof(idempotence_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/0507cbfdc01b45e59e5ee29c6e40d839.E_input', idempotence_k2_xboole_0)).
fof(t11_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k2_xboole_0(X1,X2),X3)=>r1_tarski(X1,X3)), file('data/MPTP/proofs/0507cbfdc01b45e59e5ee29c6e40d839.E_input', t11_xboole_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/0507cbfdc01b45e59e5ee29c6e40d839.E_input', commutativity_k2_tarski)).
fof(t120_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&k2_zfmisc_1(X3,k2_xboole_0(X1,X2))=k2_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))), file('data/MPTP/proofs/0507cbfdc01b45e59e5ee29c6e40d839.E_input', t120_zfmisc_1)).
fof(t13_xboole_1, axiom, ![X1, X2, X3, X4]:((r1_tarski(X1,X2)&r1_tarski(X3,X4))=>r1_tarski(k2_xboole_0(X1,X3),k2_xboole_0(X2,X4))), file('data/MPTP/proofs/0507cbfdc01b45e59e5ee29c6e40d839.E_input', t13_xboole_1)).
fof(c_0_18, plain, ![X7, X8]:k3_tarski(k2_tarski(X7,X8))=k2_xboole_0(X7,X8), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_19, plain, ![X9, X10]:k1_enumset1(X9,X9,X10)=k2_tarski(X9,X10), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_20, plain, ![X11, X12]:(((r1_tarski(X11,X12)|X11!=X12)&(r1_tarski(X12,X11)|X11!=X12))&(~r1_tarski(X11,X12)|~r1_tarski(X12,X11)|X11=X12)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_21, plain, ![X13, X14, X15]:(~r1_tarski(X13,k2_xboole_0(X14,X15))|r1_tarski(k4_xboole_0(X13,X14),X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
cnf(c_0_22, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_24, plain, ![X16, X17]:r1_tarski(X16,k2_xboole_0(X16,X17)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_25, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6]:((r1_tarski(X1,k2_zfmisc_1(X3,X4))&r1_tarski(X2,k2_zfmisc_1(X5,X6)))=>r1_tarski(k2_xboole_0(X1,X2),k2_zfmisc_1(k2_xboole_0(X3,X5),k2_xboole_0(X4,X6))))), inference(assume_negation,[status(cth)],[t143_zfmisc_1])).
fof(c_0_26, plain, ![X24, X25, X26]:((r1_tarski(X24,X25)|~r1_tarski(X24,k4_xboole_0(X25,X26)))&(r1_xboole_0(X24,X26)|~r1_tarski(X24,k4_xboole_0(X25,X26)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t106_xboole_1])])])).
cnf(c_0_27, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_28, plain, ![X27]:r1_tarski(k1_xboole_0,X27), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_29, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_31, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_32, plain, ![X28, X29, X30]:(~r1_tarski(X28,X29)|~r1_xboole_0(X29,X30)|r1_xboole_0(X28,X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_xboole_1])])).
fof(c_0_33, negated_conjecture, ((r1_tarski(esk1_0,k2_zfmisc_1(esk3_0,esk4_0))&r1_tarski(esk2_0,k2_zfmisc_1(esk5_0,esk6_0)))&~r1_tarski(k2_xboole_0(esk1_0,esk2_0),k2_zfmisc_1(k2_xboole_0(esk3_0,esk5_0),k2_xboole_0(esk4_0,esk6_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])).
fof(c_0_34, plain, ![X31, X32]:(~r1_xboole_0(X31,X32)|r1_xboole_0(X32,X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
cnf(c_0_35, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k4_xboole_0(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_38, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_39, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k3_tarski(k1_enumset1(X2,X2,X3)))), inference(rw,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_40, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[c_0_31, c_0_30])).
cnf(c_0_41, plain, (r1_xboole_0(X1,X3)|~r1_tarski(X1,X2)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_42, negated_conjecture, (r1_tarski(esk2_0,k2_zfmisc_1(esk5_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_43, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_44, plain, (r1_xboole_0(k4_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
fof(c_0_45, plain, ![X33, X34, X35]:(~r1_tarski(X33,X34)|~r1_xboole_0(X33,X35)|r1_tarski(X33,k4_xboole_0(X34,X35))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t86_xboole_1])])).
cnf(c_0_46, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_47, plain, (r1_tarski(k4_xboole_0(X1,X1),X2)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_48, negated_conjecture, (r1_xboole_0(esk2_0,X1)|~r1_xboole_0(k2_zfmisc_1(esk5_0,esk6_0),X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_49, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_50, plain, (r1_tarski(X1,k4_xboole_0(X2,X3))|~r1_tarski(X1,X2)|~r1_xboole_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_51, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_52, negated_conjecture, (r1_xboole_0(esk2_0,k4_xboole_0(X1,k2_zfmisc_1(esk5_0,esk6_0)))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_53, plain, (k4_xboole_0(X1,X2)=X3|~r1_xboole_0(X3,X2)|~r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X3,X1)), inference(spm,[status(thm)],[c_0_37, c_0_50])).
fof(c_0_54, plain, ![X36, X37]:r1_tarski(k4_xboole_0(X36,X37),X36), inference(variable_rename,[status(thm)],[t36_xboole_1])).
cnf(c_0_55, plain, (r1_tarski(X1,k1_xboole_0)|~r1_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_56, negated_conjecture, (r1_xboole_0(k4_xboole_0(X1,k2_zfmisc_1(esk5_0,esk6_0)),esk2_0)), inference(spm,[status(thm)],[c_0_43, c_0_52])).
fof(c_0_57, plain, ![X38, X39]:k2_xboole_0(X38,X39)=k5_xboole_0(X38,k4_xboole_0(X39,X38)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_58, plain, ![X40]:k2_xboole_0(X40,X40)=X40, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k2_xboole_0])])).
cnf(c_0_59, plain, (k1_xboole_0=X1|~r1_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_51]), c_0_38])])).
cnf(c_0_60, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k4_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_61, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
fof(c_0_62, plain, ![X41, X42, X43]:(~r1_tarski(k2_xboole_0(X41,X42),X43)|r1_tarski(X41,X43)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_xboole_1])])).
cnf(c_0_63, negated_conjecture, (r1_tarski(k4_xboole_0(X1,k2_zfmisc_1(esk5_0,esk6_0)),k1_xboole_0)|~r1_tarski(k4_xboole_0(X1,k2_zfmisc_1(esk5_0,esk6_0)),esk2_0)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_64, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_65, plain, (k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_58])).
fof(c_0_66, plain, ![X44, X45]:k2_tarski(X44,X45)=k2_tarski(X45,X44), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_67, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(k4_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_59, c_0_44])).
cnf(c_0_68, plain, (r1_tarski(k4_xboole_0(k4_xboole_0(X1,X2),X3),X1)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_69, negated_conjecture, (r1_tarski(esk1_0,k2_zfmisc_1(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_70, plain, (r1_tarski(X1,X3)|~r1_tarski(k2_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_71, negated_conjecture, (k4_xboole_0(X1,k2_zfmisc_1(esk5_0,esk6_0))=k1_xboole_0|~r1_tarski(k4_xboole_0(X1,k2_zfmisc_1(esk5_0,esk6_0)),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_63]), c_0_38])])).
cnf(c_0_72, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[c_0_64, c_0_30])).
cnf(c_0_73, plain, (k3_tarski(k1_enumset1(X1,X1,X1))=X1), inference(rw,[status(thm)],[c_0_65, c_0_30])).
cnf(c_0_74, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_75, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_76, negated_conjecture, (r1_xboole_0(esk1_0,X1)|~r1_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),X1)), inference(spm,[status(thm)],[c_0_41, c_0_69])).
cnf(c_0_77, plain, (r1_tarski(X1,X3)|~r1_tarski(k3_tarski(k1_enumset1(X1,X1,X2)),X3)), inference(rw,[status(thm)],[c_0_70, c_0_30])).
cnf(c_0_78, negated_conjecture, (k4_xboole_0(esk2_0,k2_zfmisc_1(esk5_0,esk6_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_71, c_0_61])).
cnf(c_0_79, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_51]), c_0_73])).
cnf(c_0_80, plain, (k1_enumset1(X1,X1,X2)=k1_enumset1(X2,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_23]), c_0_23])).
fof(c_0_81, plain, ![X46, X47, X48]:(k2_zfmisc_1(k2_xboole_0(X46,X47),X48)=k2_xboole_0(k2_zfmisc_1(X46,X48),k2_zfmisc_1(X47,X48))&k2_zfmisc_1(X48,k2_xboole_0(X46,X47))=k2_xboole_0(k2_zfmisc_1(X48,X46),k2_zfmisc_1(X48,X47))), inference(variable_rename,[status(thm)],[t120_zfmisc_1])).
cnf(c_0_82, plain, (k1_xboole_0=X1|~r1_xboole_0(X1,X2)|~r1_tarski(X1,k4_xboole_0(X2,X3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_75]), c_0_38])])).
cnf(c_0_83, negated_conjecture, (r1_xboole_0(esk1_0,k4_xboole_0(X1,k2_zfmisc_1(esk3_0,esk4_0)))), inference(spm,[status(thm)],[c_0_76, c_0_49])).
fof(c_0_84, plain, ![X49, X50, X51, X52]:(~r1_tarski(X49,X50)|~r1_tarski(X51,X52)|r1_tarski(k2_xboole_0(X49,X51),k2_xboole_0(X50,X52))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_xboole_1])])).
cnf(c_0_85, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(k3_tarski(k1_enumset1(X1,X1,X2)),k3_tarski(k1_enumset1(X1,X1,X2)),X3)))), inference(spm,[status(thm)],[c_0_77, c_0_40])).
cnf(c_0_86, negated_conjecture, (k3_tarski(k1_enumset1(esk2_0,esk2_0,k2_zfmisc_1(esk5_0,esk6_0)))=k2_zfmisc_1(esk5_0,esk6_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_78]), c_0_79]), c_0_80])).
cnf(c_0_87, plain, (k2_zfmisc_1(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_88, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(k4_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_82, c_0_36])).
cnf(c_0_89, negated_conjecture, (r1_xboole_0(k4_xboole_0(X1,k2_zfmisc_1(esk3_0,esk4_0)),esk1_0)), inference(spm,[status(thm)],[c_0_43, c_0_83])).
cnf(c_0_90, plain, (r1_tarski(k2_xboole_0(X1,X3),k2_xboole_0(X2,X4))|~r1_tarski(X1,X2)|~r1_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_91, plain, (k2_zfmisc_1(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_92, negated_conjecture, (r1_tarski(esk2_0,k3_tarski(k1_enumset1(k2_zfmisc_1(esk5_0,esk6_0),k2_zfmisc_1(esk5_0,esk6_0),X1)))), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_93, plain, (k2_zfmisc_1(k3_tarski(k1_enumset1(X1,X1,X2)),X3)=k3_tarski(k1_enumset1(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_30]), c_0_30])).
cnf(c_0_94, negated_conjecture, (k4_xboole_0(esk1_0,k2_zfmisc_1(esk3_0,esk4_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_95, negated_conjecture, (~r1_tarski(k2_xboole_0(esk1_0,esk2_0),k2_zfmisc_1(k2_xboole_0(esk3_0,esk5_0),k2_xboole_0(esk4_0,esk6_0)))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_96, plain, (r1_tarski(k3_tarski(k1_enumset1(X1,X1,X3)),k3_tarski(k1_enumset1(X2,X2,X4)))|~r1_tarski(X3,X4)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_30]), c_0_30])).
cnf(c_0_97, plain, (k2_zfmisc_1(X1,k3_tarski(k1_enumset1(X2,X2,X3)))=k3_tarski(k1_enumset1(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_30]), c_0_30])).
cnf(c_0_98, negated_conjecture, (r1_tarski(esk2_0,k2_zfmisc_1(k3_tarski(k1_enumset1(esk5_0,esk5_0,X1)),esk6_0))), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_99, negated_conjecture, (k3_tarski(k1_enumset1(esk1_0,esk1_0,k2_zfmisc_1(esk3_0,esk4_0)))=k2_zfmisc_1(esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_94]), c_0_79]), c_0_80])).
cnf(c_0_100, negated_conjecture, (~r1_tarski(k3_tarski(k1_enumset1(esk1_0,esk1_0,esk2_0)),k2_zfmisc_1(k3_tarski(k1_enumset1(esk3_0,esk3_0,esk5_0)),k3_tarski(k1_enumset1(esk4_0,esk4_0,esk6_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_95, c_0_30]), c_0_30]), c_0_30])).
cnf(c_0_101, plain, (r1_tarski(k3_tarski(k1_enumset1(X1,X1,X2)),k2_zfmisc_1(X3,k3_tarski(k1_enumset1(X4,X4,X5))))|~r1_tarski(X2,k2_zfmisc_1(X3,X5))|~r1_tarski(X1,k2_zfmisc_1(X3,X4))), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_102, negated_conjecture, (r1_tarski(esk2_0,k2_zfmisc_1(k3_tarski(k1_enumset1(X1,X1,esk5_0)),esk6_0))), inference(spm,[status(thm)],[c_0_98, c_0_80])).
cnf(c_0_103, negated_conjecture, (r1_tarski(esk1_0,X1)|~r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),X1)), inference(spm,[status(thm)],[c_0_77, c_0_99])).
cnf(c_0_104, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(k3_tarski(k1_enumset1(X1,X1,X3)),X2))), inference(spm,[status(thm)],[c_0_40, c_0_93])).
cnf(c_0_105, negated_conjecture, (~r1_tarski(esk1_0,k2_zfmisc_1(k3_tarski(k1_enumset1(esk3_0,esk3_0,esk5_0)),esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_102])])).
cnf(c_0_106, negated_conjecture, (r1_tarski(esk1_0,k2_zfmisc_1(k3_tarski(k1_enumset1(esk3_0,esk3_0,X1)),esk4_0))), inference(spm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_107, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_105, c_0_106])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 108
# Proof object clause steps            : 71
# Proof object formula steps           : 37
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 35
# Proof object simplifying inferences  : 33
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 22
# Processed clauses                    : 497
# ...of these trivial                  : 14
# ...subsumed                          : 262
# ...remaining for further processing  : 221
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 13
# Backward-rewritten                   : 3
# Generated clauses                    : 1987
# ...of the previous two non-trivial   : 1450
# Contextual simplify-reflections      : 1
# Paramodulations                      : 1985
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 182
#    Positive orientable unit clauses  : 54
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 123
# Current number of unprocessed clauses: 971
# ...number of literals in the above   : 2684
# Current number of archived formulas  : 0
# Current number of archived clauses   : 39
# Clause-clause subsumption calls (NU) : 2359
# Rec. Clause-clause subsumption calls : 2064
# Non-unit clause-clause subsumptions  : 263
# Unit Clause-clause subsumption calls : 79
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 94
# BW rewrite match successes           : 31
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 22317

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.003 s
# Total time               : 0.024 s
# Maximum resident set size: 28612 pages
