# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/05176bef7dcc4551877cb4cd52a50ff9.E_input', d1_xboole_0)).
fof(t6_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=>~((r2_hidden(X1,X4)&![X5, X6]:~((X1=k4_tarski(X5,X6)&(r2_hidden(X5,X2)&r2_hidden(X6,X3))))))), file('data/MPTP/proofs/05176bef7dcc4551877cb4cd52a50ff9.E_input', t6_relset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/05176bef7dcc4551877cb4cd52a50ff9.E_input', t3_subset)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/05176bef7dcc4551877cb4cd52a50ff9.E_input', d10_xboole_0)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/05176bef7dcc4551877cb4cd52a50ff9.E_input', t7_xboole_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/05176bef7dcc4551877cb4cd52a50ff9.E_input', fc1_xboole_0)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/05176bef7dcc4551877cb4cd52a50ff9.E_input', d4_relat_1)).
fof(d2_funct_2, axiom, ![X1, X2, X3]:(X3=k1_funct_2(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:((v1_relat_1(X5)&v1_funct_1(X5))&(X4=X5&(k1_relat_1(X5)=X1&r1_tarski(k2_relat_1(X5),X2)))))), file('data/MPTP/proofs/05176bef7dcc4551877cb4cd52a50ff9.E_input', d2_funct_2)).
fof(t121_funct_2, conjecture, ![X1, X2, X3]:(r2_hidden(X3,k1_funct_2(X1,X2))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))), file('data/MPTP/proofs/05176bef7dcc4551877cb4cd52a50ff9.E_input', t121_funct_2)).
fof(t3_funct_2, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_funct_1(X1)&(v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))))), file('data/MPTP/proofs/05176bef7dcc4551877cb4cd52a50ff9.E_input', t3_funct_2)).
fof(t11_relset_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>((r1_tarski(k1_relat_1(X3),X1)&r1_tarski(k2_relat_1(X3),X2))=>m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))), file('data/MPTP/proofs/05176bef7dcc4551877cb4cd52a50ff9.E_input', t11_relset_1)).
fof(cc5_funct_2, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_funct_2(X3,X1,X2))=>(v1_funct_1(X3)&v1_partfun1(X3,X1))))), file('data/MPTP/proofs/05176bef7dcc4551877cb4cd52a50ff9.E_input', cc5_funct_2)).
fof(t5_funct_2, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X3)=X1&![X4]:(r2_hidden(X4,X1)=>r2_hidden(k1_funct_1(X3,X4),X2)))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))), file('data/MPTP/proofs/05176bef7dcc4551877cb4cd52a50ff9.E_input', t5_funct_2)).
fof(cc1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_partfun1(X3,X1))=>(v1_funct_1(X3)&v1_funct_2(X3,X1,X2)))), file('data/MPTP/proofs/05176bef7dcc4551877cb4cd52a50ff9.E_input', cc1_funct_2)).
fof(c_0_14, plain, ![X7, X8, X9]:((~v1_xboole_0(X7)|~r2_hidden(X8,X7))&(r2_hidden(esk1_1(X9),X9)|v1_xboole_0(X9))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_15, plain, ![X40, X41, X42, X43]:((X40=k4_tarski(esk12_4(X40,X41,X42,X43),esk13_4(X40,X41,X42,X43))|~r2_hidden(X40,X43)|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(X41,X42))))&((r2_hidden(esk12_4(X40,X41,X42,X43),X41)|~r2_hidden(X40,X43)|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(X41,X42))))&(r2_hidden(esk13_4(X40,X41,X42,X43),X42)|~r2_hidden(X40,X43)|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(X41,X42)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_relset_1])])])])).
cnf(c_0_16, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_17, plain, (r2_hidden(esk12_4(X1,X2,X3,X4),X2)|~r2_hidden(X1,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_18, plain, ![X46, X47]:((~m1_subset_1(X46,k1_zfmisc_1(X47))|r1_tarski(X46,X47))&(~r1_tarski(X46,X47)|m1_subset_1(X46,k1_zfmisc_1(X47)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_19, plain, ![X48, X49]:(((r1_tarski(X48,X49)|X48!=X49)&(r1_tarski(X49,X48)|X48!=X49))&(~r1_tarski(X48,X49)|~r1_tarski(X49,X48)|X48=X49)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_20, plain, (~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r2_hidden(X4,X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_21, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23, plain, (~r1_tarski(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X4,X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_24, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_22])).
fof(c_0_25, plain, ![X38]:(X38=k1_xboole_0|r2_hidden(esk11_1(X38),X38)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_26, plain, (~r2_hidden(X1,k2_zfmisc_1(X2,X3))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_27, plain, (X1=k1_xboole_0|r2_hidden(esk11_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_28, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_29, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_30, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_28, c_0_29])).
fof(c_0_31, plain, ![X11, X12, X13, X15, X16, X17, X18, X20]:(((~r2_hidden(X13,X12)|r2_hidden(k4_tarski(X13,esk2_3(X11,X12,X13)),X11)|X12!=k1_relat_1(X11))&(~r2_hidden(k4_tarski(X15,X16),X11)|r2_hidden(X15,X12)|X12!=k1_relat_1(X11)))&((~r2_hidden(esk3_2(X17,X18),X18)|~r2_hidden(k4_tarski(esk3_2(X17,X18),X20),X17)|X18=k1_relat_1(X17))&(r2_hidden(esk3_2(X17,X18),X18)|r2_hidden(k4_tarski(esk3_2(X17,X18),esk4_2(X17,X18)),X17)|X18=k1_relat_1(X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
cnf(c_0_32, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_30]), c_0_29])])).
cnf(c_0_33, plain, (r2_hidden(k4_tarski(X1,esk2_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_34, plain, (X1!=k1_relat_1(k1_xboole_0)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_35, plain, (~r2_hidden(X1,k1_relat_1(k1_xboole_0))), inference(er,[status(thm)],[c_0_34])).
fof(c_0_36, plain, ![X22, X23, X24, X25, X27, X28, X29, X30, X31, X33]:(((((v1_relat_1(esk5_4(X22,X23,X24,X25))|~r2_hidden(X25,X24)|X24!=k1_funct_2(X22,X23))&(v1_funct_1(esk5_4(X22,X23,X24,X25))|~r2_hidden(X25,X24)|X24!=k1_funct_2(X22,X23)))&((X25=esk5_4(X22,X23,X24,X25)|~r2_hidden(X25,X24)|X24!=k1_funct_2(X22,X23))&((k1_relat_1(esk5_4(X22,X23,X24,X25))=X22|~r2_hidden(X25,X24)|X24!=k1_funct_2(X22,X23))&(r1_tarski(k2_relat_1(esk5_4(X22,X23,X24,X25)),X23)|~r2_hidden(X25,X24)|X24!=k1_funct_2(X22,X23)))))&(~v1_relat_1(X28)|~v1_funct_1(X28)|(X27!=X28|(k1_relat_1(X28)!=X22|~r1_tarski(k2_relat_1(X28),X23)))|r2_hidden(X27,X24)|X24!=k1_funct_2(X22,X23)))&((~r2_hidden(esk6_3(X29,X30,X31),X31)|(~v1_relat_1(X33)|~v1_funct_1(X33)|(esk6_3(X29,X30,X31)!=X33|(k1_relat_1(X33)!=X29|~r1_tarski(k2_relat_1(X33),X30))))|X31=k1_funct_2(X29,X30))&(((v1_relat_1(esk7_3(X29,X30,X31))|r2_hidden(esk6_3(X29,X30,X31),X31)|X31=k1_funct_2(X29,X30))&(v1_funct_1(esk7_3(X29,X30,X31))|r2_hidden(esk6_3(X29,X30,X31),X31)|X31=k1_funct_2(X29,X30)))&((esk6_3(X29,X30,X31)=esk7_3(X29,X30,X31)|r2_hidden(esk6_3(X29,X30,X31),X31)|X31=k1_funct_2(X29,X30))&((k1_relat_1(esk7_3(X29,X30,X31))=X29|r2_hidden(esk6_3(X29,X30,X31),X31)|X31=k1_funct_2(X29,X30))&(r1_tarski(k2_relat_1(esk7_3(X29,X30,X31)),X30)|r2_hidden(esk6_3(X29,X30,X31),X31)|X31=k1_funct_2(X29,X30))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_funct_2])])])])])])).
cnf(c_0_37, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_35, c_0_27])).
cnf(c_0_38, plain, (k1_relat_1(esk5_4(X1,X2,X3,X4))=X1|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_39, plain, (X1=esk5_4(X2,X3,X4,X1)|~r2_hidden(X1,X4)|X4!=k1_funct_2(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_40, negated_conjecture, ~(![X1, X2, X3]:(r2_hidden(X3,k1_funct_2(X1,X2))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))), inference(assume_negation,[status(cth)],[t121_funct_2])).
cnf(c_0_41, plain, (v1_funct_1(esk5_4(X1,X2,X3,X4))|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_42, plain, (v1_relat_1(esk5_4(X1,X2,X3,X4))|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_43, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(rw,[status(thm)],[c_0_34, c_0_37])).
cnf(c_0_44, plain, (k1_relat_1(X1)=X2|X3!=k1_funct_2(X2,X4)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
fof(c_0_45, negated_conjecture, (r2_hidden(esk10_0,k1_funct_2(esk8_0,esk9_0))&(~v1_funct_1(esk10_0)|(~v1_funct_2(esk10_0,esk8_0,esk9_0)|~m1_subset_1(esk10_0,k1_zfmisc_1(k2_zfmisc_1(esk8_0,esk9_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_40])])])).
cnf(c_0_46, plain, (v1_funct_1(X1)|X2!=k1_funct_2(X3,X4)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_41, c_0_39])).
cnf(c_0_47, plain, (v1_relat_1(X1)|X2!=k1_funct_2(X3,X4)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_42, c_0_39])).
cnf(c_0_48, plain, (X1!=k1_relat_1(X2)|X2!=k1_xboole_0|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_43, c_0_33])).
fof(c_0_49, plain, ![X50]:((v1_funct_1(X50)|(~v1_relat_1(X50)|~v1_funct_1(X50)))&((v1_funct_2(X50,k1_relat_1(X50),k2_relat_1(X50))|(~v1_relat_1(X50)|~v1_funct_1(X50)))&(m1_subset_1(X50,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X50),k2_relat_1(X50))))|(~v1_relat_1(X50)|~v1_funct_1(X50))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_funct_2])])])).
cnf(c_0_50, plain, (k1_relat_1(X1)=X2|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_51, negated_conjecture, (r2_hidden(esk10_0,k1_funct_2(esk8_0,esk9_0))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_52, plain, (v1_funct_1(X1)|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(er,[status(thm)],[c_0_46])).
cnf(c_0_53, plain, (v1_relat_1(X1)|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(er,[status(thm)],[c_0_47])).
cnf(c_0_54, plain, (X1!=k1_xboole_0|~r2_hidden(X2,k1_relat_1(X1))), inference(er,[status(thm)],[c_0_48])).
cnf(c_0_55, plain, (r2_hidden(esk1_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_56, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_57, negated_conjecture, (k1_relat_1(esk10_0)=esk8_0), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_58, negated_conjecture, (v1_funct_1(esk10_0)), inference(spm,[status(thm)],[c_0_52, c_0_51])).
cnf(c_0_59, negated_conjecture, (v1_relat_1(esk10_0)), inference(spm,[status(thm)],[c_0_53, c_0_51])).
cnf(c_0_60, plain, (X1!=k1_relat_1(X2)|~r2_hidden(X3,X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_16, c_0_33])).
fof(c_0_61, plain, ![X58, X59, X60]:(~v1_relat_1(X60)|(~r1_tarski(k1_relat_1(X60),X58)|~r1_tarski(k2_relat_1(X60),X59)|m1_subset_1(X60,k1_zfmisc_1(k2_zfmisc_1(X58,X59))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_relset_1])])).
cnf(c_0_62, plain, (r2_hidden(esk13_4(X1,X2,X3,X4),X3)|~r2_hidden(X1,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_63, plain, ![X51, X52, X53]:((v1_funct_1(X53)|(~v1_funct_1(X53)|~v1_funct_2(X53,X51,X52))|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(X51,X52)))|v1_xboole_0(X52))&(v1_partfun1(X53,X51)|(~v1_funct_1(X53)|~v1_funct_2(X53,X51,X52))|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(X51,X52)))|v1_xboole_0(X52))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_funct_2])])])])])).
cnf(c_0_64, plain, (v1_xboole_0(k1_relat_1(X1))|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_65, negated_conjecture, (m1_subset_1(esk10_0,k1_zfmisc_1(k2_zfmisc_1(esk8_0,k2_relat_1(esk10_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58]), c_0_59])])).
cnf(c_0_66, plain, (~r2_hidden(X1,k1_relat_1(X2))|~v1_xboole_0(X2)), inference(er,[status(thm)],[c_0_60])).
cnf(c_0_67, negated_conjecture, (~v1_funct_1(esk10_0)|~v1_funct_2(esk10_0,esk8_0,esk9_0)|~m1_subset_1(esk10_0,k1_zfmisc_1(k2_zfmisc_1(esk8_0,esk9_0)))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_68, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_69, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_70, plain, (r1_tarski(k2_relat_1(esk5_4(X1,X2,X3,X4)),X2)|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_71, plain, (~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r2_hidden(X4,X1)|~v1_xboole_0(X3)), inference(spm,[status(thm)],[c_0_16, c_0_62])).
cnf(c_0_72, plain, (v1_partfun1(X1,X2)|v1_xboole_0(X3)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_73, plain, (v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_74, plain, (k2_zfmisc_1(k1_relat_1(X1),X2)=k1_xboole_0|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_28, c_0_64])).
cnf(c_0_75, negated_conjecture, (m1_subset_1(esk10_0,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(esk8_0)), inference(spm,[status(thm)],[c_0_65, c_0_28])).
cnf(c_0_76, negated_conjecture, (v1_xboole_0(esk8_0)|k1_xboole_0!=esk10_0), inference(spm,[status(thm)],[c_0_64, c_0_57])).
fof(c_0_77, plain, ![X54, X55, X56]:(((v1_funct_1(X56)|(r2_hidden(esk14_3(X54,X55,X56),X54)|k1_relat_1(X56)!=X54)|(~v1_relat_1(X56)|~v1_funct_1(X56)))&((v1_funct_2(X56,X54,X55)|(r2_hidden(esk14_3(X54,X55,X56),X54)|k1_relat_1(X56)!=X54)|(~v1_relat_1(X56)|~v1_funct_1(X56)))&(m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(X54,X55)))|(r2_hidden(esk14_3(X54,X55,X56),X54)|k1_relat_1(X56)!=X54)|(~v1_relat_1(X56)|~v1_funct_1(X56)))))&((v1_funct_1(X56)|(~r2_hidden(k1_funct_1(X56,esk14_3(X54,X55,X56)),X55)|k1_relat_1(X56)!=X54)|(~v1_relat_1(X56)|~v1_funct_1(X56)))&((v1_funct_2(X56,X54,X55)|(~r2_hidden(k1_funct_1(X56,esk14_3(X54,X55,X56)),X55)|k1_relat_1(X56)!=X54)|(~v1_relat_1(X56)|~v1_funct_1(X56)))&(m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(X54,X55)))|(~r2_hidden(k1_funct_1(X56,esk14_3(X54,X55,X56)),X55)|k1_relat_1(X56)!=X54)|(~v1_relat_1(X56)|~v1_funct_1(X56)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_funct_2])])])])).
cnf(c_0_78, plain, (k1_relat_1(X1)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_66, c_0_27])).
cnf(c_0_79, plain, (v1_xboole_0(X1)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_43, c_0_55])).
cnf(c_0_80, negated_conjecture, (~m1_subset_1(esk10_0,k1_zfmisc_1(k2_zfmisc_1(esk8_0,esk9_0)))|~v1_funct_2(esk10_0,esk8_0,esk9_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_58])])).
cnf(c_0_81, plain, (r1_tarski(X1,k2_zfmisc_1(X2,X3))|~r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X3)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_82, plain, (r1_tarski(k2_relat_1(X1),X2)|X3!=k1_funct_2(X4,X2)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_70, c_0_39])).
cnf(c_0_83, negated_conjecture, (~r2_hidden(X1,esk10_0)|~v1_xboole_0(k2_relat_1(esk10_0))), inference(spm,[status(thm)],[c_0_71, c_0_65])).
cnf(c_0_84, plain, (v1_partfun1(X1,k1_relat_1(X1))|v1_xboole_0(k2_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_56]), c_0_73])).
cnf(c_0_85, negated_conjecture, (k2_zfmisc_1(esk8_0,X1)=k1_xboole_0|k1_xboole_0!=esk10_0), inference(spm,[status(thm)],[c_0_74, c_0_57])).
cnf(c_0_86, negated_conjecture, (m1_subset_1(esk10_0,k1_zfmisc_1(k1_xboole_0))|k1_xboole_0!=esk10_0), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_87, plain, (v1_funct_2(X1,X2,X3)|r2_hidden(esk14_3(X2,X3,X1),X2)|k1_relat_1(X1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_88, plain, (k1_relat_1(X1)=k1_xboole_0|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_89, negated_conjecture, (~v1_funct_2(esk10_0,esk8_0,esk9_0)|~r1_tarski(k1_relat_1(esk10_0),esk8_0)|~r1_tarski(k2_relat_1(esk10_0),esk9_0)|~v1_relat_1(esk10_0)), inference(spm,[status(thm)],[c_0_80, c_0_69])).
fof(c_0_90, plain, ![X61, X62, X63]:((v1_funct_1(X63)|(~v1_funct_1(X63)|~v1_partfun1(X63,X61))|~m1_subset_1(X63,k1_zfmisc_1(k2_zfmisc_1(X61,X62))))&(v1_funct_2(X63,X61,X62)|(~v1_funct_1(X63)|~v1_partfun1(X63,X61))|~m1_subset_1(X63,k1_zfmisc_1(k2_zfmisc_1(X61,X62))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_2])])])).
cnf(c_0_91, negated_conjecture, (r1_tarski(esk10_0,k2_zfmisc_1(X1,X2))|~r1_tarski(k2_relat_1(esk10_0),X2)|~r1_tarski(esk8_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_57]), c_0_59])])).
cnf(c_0_92, plain, (r1_tarski(k2_relat_1(X1),X2)|~r2_hidden(X1,k1_funct_2(X3,X2))), inference(er,[status(thm)],[c_0_82])).
cnf(c_0_93, negated_conjecture, (v1_partfun1(esk10_0,esk8_0)|~r2_hidden(X1,esk10_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_57]), c_0_58]), c_0_59])])).
cnf(c_0_94, negated_conjecture, (k1_xboole_0!=esk10_0|~v1_funct_2(esk10_0,esk8_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_85]), c_0_86])).
cnf(c_0_95, plain, (v1_funct_2(X1,X2,X3)|k1_xboole_0!=X2|X1!=k1_xboole_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_43])).
cnf(c_0_96, negated_conjecture, (k1_xboole_0=esk8_0|k1_xboole_0!=esk10_0), inference(spm,[status(thm)],[c_0_88, c_0_57])).
cnf(c_0_97, negated_conjecture, (~v1_funct_2(esk10_0,esk8_0,esk9_0)|~r1_tarski(k1_relat_1(esk10_0),esk8_0)|~r1_tarski(k2_relat_1(esk10_0),esk9_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_59])])).
cnf(c_0_98, plain, (v1_funct_2(X1,X2,X3)|~v1_funct_1(X1)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_99, negated_conjecture, (r1_tarski(esk10_0,k2_zfmisc_1(esk8_0,X1))|~r1_tarski(k2_relat_1(esk10_0),X1)), inference(spm,[status(thm)],[c_0_91, c_0_24])).
cnf(c_0_100, negated_conjecture, (r1_tarski(k2_relat_1(esk10_0),esk9_0)), inference(spm,[status(thm)],[c_0_92, c_0_51])).
cnf(c_0_101, negated_conjecture, (k1_xboole_0=esk10_0|v1_partfun1(esk10_0,esk8_0)), inference(spm,[status(thm)],[c_0_93, c_0_27])).
cnf(c_0_102, negated_conjecture, (k1_xboole_0!=esk10_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_58]), c_0_59])]), c_0_96])).
cnf(c_0_103, negated_conjecture, (~v1_funct_2(esk10_0,esk8_0,esk9_0)|~r1_tarski(k2_relat_1(esk10_0),esk9_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_57]), c_0_24])])).
cnf(c_0_104, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_98, c_0_21])).
cnf(c_0_105, negated_conjecture, (r1_tarski(esk10_0,k2_zfmisc_1(esk8_0,esk9_0))), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_106, negated_conjecture, (v1_partfun1(esk10_0,esk8_0)), inference(sr,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_107, negated_conjecture, (~v1_funct_2(esk10_0,esk8_0,esk9_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_100])])).
cnf(c_0_108, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_106]), c_0_58])]), c_0_107]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 109
# Proof object clause steps            : 81
# Proof object formula steps           : 28
# Proof object conjectures             : 30
# Proof object clause conjectures      : 27
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 51
# Proof object simplifying inferences  : 34
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 44
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 39
# Processed clauses                    : 4743
# ...of these trivial                  : 26
# ...subsumed                          : 3846
# ...remaining for further processing  : 871
# Other redundant clauses eliminated   : 166
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 36
# Backward-rewritten                   : 20
# Generated clauses                    : 23787
# ...of the previous two non-trivial   : 17395
# Contextual simplify-reflections      : 56
# Paramodulations                      : 23475
# Factorizations                       : 2
# Equation resolutions                 : 309
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 773
#    Positive orientable unit clauses  : 18
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 11
#    Non-unit-clauses                  : 744
# Current number of unprocessed clauses: 12637
# ...number of literals in the above   : 60482
# Current number of archived formulas  : 0
# Current number of archived clauses   : 95
# Clause-clause subsumption calls (NU) : 106707
# Rec. Clause-clause subsumption calls : 44783
# Non-unit clause-clause subsumptions  : 1752
# Unit Clause-clause subsumption calls : 1698
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 298290

# -------------------------------------------------
# User time                : 0.198 s
# System time              : 0.023 s
# Total time               : 0.221 s
# Maximum resident set size: 28660 pages
