# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d12_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(X3=k8_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>(r2_hidden(X5,X1)&r2_hidden(k4_tarski(X4,X5),X2)))))), file('data/MPTP/proofs/0519ace3c5234e3a8b13010aa5bf19a1.E_input', d12_relat_1)).
fof(dt_k8_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>v1_relat_1(k8_relat_1(X1,X2))), file('data/MPTP/proofs/0519ace3c5234e3a8b13010aa5bf19a1.E_input', dt_k8_relat_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/0519ace3c5234e3a8b13010aa5bf19a1.E_input', d5_relat_1)).
fof(t115_relat_1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k2_relat_1(k8_relat_1(X2,X3)))<=>(r2_hidden(X1,X2)&r2_hidden(X1,k2_relat_1(X3))))), file('data/MPTP/proofs/0519ace3c5234e3a8b13010aa5bf19a1.E_input', t115_relat_1)).
fof(c_0_4, plain, ![X6, X7, X8, X9, X10, X11, X12]:((((r2_hidden(X10,X6)|~r2_hidden(k4_tarski(X9,X10),X8)|X8!=k8_relat_1(X6,X7)|~v1_relat_1(X8)|~v1_relat_1(X7))&(r2_hidden(k4_tarski(X9,X10),X7)|~r2_hidden(k4_tarski(X9,X10),X8)|X8!=k8_relat_1(X6,X7)|~v1_relat_1(X8)|~v1_relat_1(X7)))&(~r2_hidden(X12,X6)|~r2_hidden(k4_tarski(X11,X12),X7)|r2_hidden(k4_tarski(X11,X12),X8)|X8!=k8_relat_1(X6,X7)|~v1_relat_1(X8)|~v1_relat_1(X7)))&((~r2_hidden(k4_tarski(esk1_3(X6,X7,X8),esk2_3(X6,X7,X8)),X8)|(~r2_hidden(esk2_3(X6,X7,X8),X6)|~r2_hidden(k4_tarski(esk1_3(X6,X7,X8),esk2_3(X6,X7,X8)),X7))|X8=k8_relat_1(X6,X7)|~v1_relat_1(X8)|~v1_relat_1(X7))&((r2_hidden(esk2_3(X6,X7,X8),X6)|r2_hidden(k4_tarski(esk1_3(X6,X7,X8),esk2_3(X6,X7,X8)),X8)|X8=k8_relat_1(X6,X7)|~v1_relat_1(X8)|~v1_relat_1(X7))&(r2_hidden(k4_tarski(esk1_3(X6,X7,X8),esk2_3(X6,X7,X8)),X7)|r2_hidden(k4_tarski(esk1_3(X6,X7,X8),esk2_3(X6,X7,X8)),X8)|X8=k8_relat_1(X6,X7)|~v1_relat_1(X8)|~v1_relat_1(X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_relat_1])])])])])])).
fof(c_0_5, plain, ![X15, X16]:(~v1_relat_1(X16)|v1_relat_1(k8_relat_1(X15,X16))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_relat_1])])).
fof(c_0_6, plain, ![X17, X18, X19, X21, X22, X23, X24, X26]:(((~r2_hidden(X19,X18)|r2_hidden(k4_tarski(esk3_3(X17,X18,X19),X19),X17)|X18!=k2_relat_1(X17))&(~r2_hidden(k4_tarski(X22,X21),X17)|r2_hidden(X21,X18)|X18!=k2_relat_1(X17)))&((~r2_hidden(esk4_2(X23,X24),X24)|~r2_hidden(k4_tarski(X26,esk4_2(X23,X24)),X23)|X24=k2_relat_1(X23))&(r2_hidden(esk4_2(X23,X24),X24)|r2_hidden(k4_tarski(esk5_2(X23,X24),esk4_2(X23,X24)),X23)|X24=k2_relat_1(X23)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_7, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),X4)|X4!=k8_relat_1(X2,X5)|~v1_relat_1(X4)|~v1_relat_1(X5)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, plain, (v1_relat_1(k8_relat_1(X2,X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (r2_hidden(k4_tarski(esk3_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_10, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k2_relat_1(k8_relat_1(X2,X3)))<=>(r2_hidden(X1,X2)&r2_hidden(X1,k2_relat_1(X3)))))), inference(assume_negation,[status(cth)],[t115_relat_1])).
cnf(c_0_11, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k8_relat_1(X2,X4))|~v1_relat_1(X4)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_7]), c_0_8])).
cnf(c_0_12, plain, (r2_hidden(k4_tarski(esk3_3(X1,k2_relat_1(X1),X2),X2),X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_9])).
fof(c_0_13, negated_conjecture, (v1_relat_1(esk8_0)&((~r2_hidden(esk6_0,k2_relat_1(k8_relat_1(esk7_0,esk8_0)))|(~r2_hidden(esk6_0,esk7_0)|~r2_hidden(esk6_0,k2_relat_1(esk8_0))))&((r2_hidden(esk6_0,esk7_0)|r2_hidden(esk6_0,k2_relat_1(k8_relat_1(esk7_0,esk8_0))))&(r2_hidden(esk6_0,k2_relat_1(esk8_0))|r2_hidden(esk6_0,k2_relat_1(k8_relat_1(esk7_0,esk8_0))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])).
cnf(c_0_14, plain, (r2_hidden(X2,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15, plain, (r2_hidden(k4_tarski(X3,X1),X5)|~r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),X4)|X5!=k8_relat_1(X2,X4)|~v1_relat_1(X5)|~v1_relat_1(X4)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_16, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k2_relat_1(k8_relat_1(X2,X3)))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_17, negated_conjecture, (r2_hidden(esk6_0,esk7_0)|r2_hidden(esk6_0,k2_relat_1(k8_relat_1(esk7_0,esk8_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, negated_conjecture, (v1_relat_1(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)), inference(er,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (r2_hidden(k4_tarski(X1,X2),k8_relat_1(X3,X4))|~r2_hidden(k4_tarski(X1,X2),X4)|~r2_hidden(X2,X3)|~v1_relat_1(X4)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_15]), c_0_8])).
cnf(c_0_21, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(k4_tarski(X1,X2),X4)|X4!=k8_relat_1(X5,X3)|~v1_relat_1(X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_22, negated_conjecture, (~r2_hidden(esk6_0,k2_relat_1(k8_relat_1(esk7_0,esk8_0)))|~r2_hidden(esk6_0,esk7_0)|~r2_hidden(esk6_0,k2_relat_1(esk8_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23, negated_conjecture, (r2_hidden(esk6_0,esk7_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])])).
cnf(c_0_24, plain, (r2_hidden(X1,k2_relat_1(k8_relat_1(X2,X3)))|~r2_hidden(k4_tarski(X4,X1),X3)|~r2_hidden(X1,X2)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_25, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(k4_tarski(X1,X2),k8_relat_1(X4,X3))|~v1_relat_1(X3)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_21]), c_0_8])).
cnf(c_0_26, negated_conjecture, (~r2_hidden(esk6_0,k2_relat_1(k8_relat_1(esk7_0,esk8_0)))|~r2_hidden(esk6_0,k2_relat_1(esk8_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_23])])).
cnf(c_0_27, plain, (r2_hidden(X1,k2_relat_1(k8_relat_1(X2,X3)))|~r2_hidden(X1,k2_relat_1(X3))|~r2_hidden(X1,X2)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_24, c_0_12])).
cnf(c_0_28, plain, (r2_hidden(k4_tarski(esk3_3(k8_relat_1(X1,X2),k2_relat_1(k8_relat_1(X1,X2)),X3),X3),X2)|~r2_hidden(X3,k2_relat_1(k8_relat_1(X1,X2)))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_25, c_0_12])).
cnf(c_0_29, negated_conjecture, (r2_hidden(esk6_0,k2_relat_1(esk8_0))|r2_hidden(esk6_0,k2_relat_1(k8_relat_1(esk7_0,esk8_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_30, negated_conjecture, (~r2_hidden(esk6_0,k2_relat_1(esk8_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_23]), c_0_18])])).
cnf(c_0_31, plain, (r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(X1,k2_relat_1(k8_relat_1(X3,X2)))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_19, c_0_28])).
cnf(c_0_32, negated_conjecture, (r2_hidden(esk6_0,k2_relat_1(k8_relat_1(esk7_0,esk8_0)))), inference(sr,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_33, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_18])]), c_0_30]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 34
# Proof object clause steps            : 25
# Proof object formula steps           : 9
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 19
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 15
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 15
# Processed clauses                    : 139
# ...of these trivial                  : 1
# ...subsumed                          : 31
# ...remaining for further processing  : 107
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 3
# Generated clauses                    : 598
# ...of the previous two non-trivial   : 536
# Contextual simplify-reflections      : 9
# Paramodulations                      : 588
# Factorizations                       : 4
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 81
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 76
# Current number of unprocessed clauses: 427
# ...number of literals in the above   : 2131
# Current number of archived formulas  : 0
# Current number of archived clauses   : 21
# Clause-clause subsumption calls (NU) : 803
# Rec. Clause-clause subsumption calls : 414
# Non-unit clause-clause subsumptions  : 41
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 21957

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.004 s
# Total time               : 0.019 s
# Maximum resident set size: 28588 pages
