# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_S0e
# and selection function SelectLargestNegLit.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_yellow_0, conjecture, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>((g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))&v3_lattice3(X1))=>v3_lattice3(X2)))), file('data/MPTP/proofs/05c98e1abacc4cecb53a401685a6d453.E_input', t3_yellow_0)).
fof(dt_u1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))), file('data/MPTP/proofs/05c98e1abacc4cecb53a401685a6d453.E_input', dt_u1_orders_2)).
fof(free_g1_orders_2, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=>![X3, X4]:(g1_orders_2(X1,X2)=g1_orders_2(X3,X4)=>(X1=X3&X2=X4))), file('data/MPTP/proofs/05c98e1abacc4cecb53a401685a6d453.E_input', free_g1_orders_2)).
fof(t2_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3, X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(X4=X5=>((r2_lattice3(X1,X3,X4)=>r2_lattice3(X2,X3,X5))&(r1_lattice3(X1,X3,X4)=>r1_lattice3(X2,X3,X5))))))))), file('data/MPTP/proofs/05c98e1abacc4cecb53a401685a6d453.E_input', t2_yellow_0)).
fof(d12_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v3_lattice3(X1)<=>![X2]:?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4))))))), file('data/MPTP/proofs/05c98e1abacc4cecb53a401685a6d453.E_input', d12_lattice3)).
fof(t1_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>((X3=X5&X4=X6)=>((r1_orders_2(X1,X3,X4)=>r1_orders_2(X2,X5,X6))&(r2_orders_2(X1,X3,X4)=>r2_orders_2(X2,X5,X6))))))))))), file('data/MPTP/proofs/05c98e1abacc4cecb53a401685a6d453.E_input', t1_yellow_0)).
fof(c_0_6, negated_conjecture, ~(![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>((g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))&v3_lattice3(X1))=>v3_lattice3(X2))))), inference(assume_negation,[status(cth)],[t3_yellow_0])).
fof(c_0_7, plain, ![X9]:(~l1_orders_2(X9)|m1_subset_1(u1_orders_2(X9),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X9))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_orders_2])])).
fof(c_0_8, negated_conjecture, (l1_orders_2(esk1_0)&(l1_orders_2(esk2_0)&((g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))&v3_lattice3(esk1_0))&~v3_lattice3(esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
fof(c_0_9, plain, ![X10, X11, X12, X13]:((X10=X12|g1_orders_2(X10,X11)!=g1_orders_2(X12,X13)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X10,X10))))&(X11=X13|g1_orders_2(X10,X11)!=g1_orders_2(X12,X13)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X10,X10))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_orders_2])])])])).
cnf(c_0_10, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, negated_conjecture, (l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (X1=X2|g1_orders_2(X3,X1)!=g1_orders_2(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X3)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, negated_conjecture, (m1_subset_1(u1_orders_2(esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_14, negated_conjecture, (u1_orders_2(esk2_0)=X1|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(X2,X1)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_15, negated_conjecture, (g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_16, plain, ![X14, X15, X16, X17, X18]:((~r2_lattice3(X14,X16,X17)|r2_lattice3(X15,X16,X18)|X17!=X18|~m1_subset_1(X18,u1_struct_0(X15))|~m1_subset_1(X17,u1_struct_0(X14))|g1_orders_2(u1_struct_0(X14),u1_orders_2(X14))!=g1_orders_2(u1_struct_0(X15),u1_orders_2(X15))|~l1_orders_2(X15)|~l1_orders_2(X14))&(~r1_lattice3(X14,X16,X17)|r1_lattice3(X15,X16,X18)|X17!=X18|~m1_subset_1(X18,u1_struct_0(X15))|~m1_subset_1(X17,u1_struct_0(X14))|g1_orders_2(u1_struct_0(X14),u1_orders_2(X14))!=g1_orders_2(u1_struct_0(X15),u1_orders_2(X15))|~l1_orders_2(X15)|~l1_orders_2(X14))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_yellow_0])])])])).
cnf(c_0_17, plain, (X1=X2|g1_orders_2(X1,X3)!=g1_orders_2(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_18, negated_conjecture, (u1_orders_2(esk1_0)=u1_orders_2(esk2_0)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_19, plain, (r2_lattice3(X4,X2,X5)|~r2_lattice3(X1,X2,X3)|X3!=X5|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))|~l1_orders_2(X4)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, negated_conjecture, (u1_struct_0(esk2_0)=X1|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(X1,X2)), inference(spm,[status(thm)],[c_0_17, c_0_13])).
cnf(c_0_21, negated_conjecture, (g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk2_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))), inference(rw,[status(thm)],[c_0_15, c_0_18])).
fof(c_0_22, plain, ![X19, X20, X22, X24]:(((m1_subset_1(esk3_2(X19,X20),u1_struct_0(X19))|~v3_lattice3(X19)|~l1_orders_2(X19))&((r2_lattice3(X19,X20,esk3_2(X19,X20))|~v3_lattice3(X19)|~l1_orders_2(X19))&(~m1_subset_1(X22,u1_struct_0(X19))|(~r2_lattice3(X19,X20,X22)|r1_orders_2(X19,esk3_2(X19,X20),X22))|~v3_lattice3(X19)|~l1_orders_2(X19))))&((m1_subset_1(esk5_2(X19,X24),u1_struct_0(X19))|~r2_lattice3(X19,esk4_1(X19),X24)|~m1_subset_1(X24,u1_struct_0(X19))|v3_lattice3(X19)|~l1_orders_2(X19))&((r2_lattice3(X19,esk4_1(X19),esk5_2(X19,X24))|~r2_lattice3(X19,esk4_1(X19),X24)|~m1_subset_1(X24,u1_struct_0(X19))|v3_lattice3(X19)|~l1_orders_2(X19))&(~r1_orders_2(X19,X24,esk5_2(X19,X24))|~r2_lattice3(X19,esk4_1(X19),X24)|~m1_subset_1(X24,u1_struct_0(X19))|v3_lattice3(X19)|~l1_orders_2(X19))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_lattice3])])])])])).
cnf(c_0_23, plain, (r2_lattice3(X1,X2,X3)|g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(X4,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X4))|~l1_orders_2(X1)|~l1_orders_2(X4)), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_24, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_25, negated_conjecture, (u1_struct_0(esk1_0)=u1_struct_0(esk2_0)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_26, plain, (m1_subset_1(esk3_2(X1,X2),u1_struct_0(X1))|~v3_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27, negated_conjecture, (v3_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_28, negated_conjecture, (r2_lattice3(X1,X2,X3)|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(esk1_0,X2,X3)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_18]), c_0_24])]), c_0_25]), c_0_25])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk3_2(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_24])])).
cnf(c_0_30, negated_conjecture, (r2_lattice3(esk2_0,X1,X2)|~r2_lattice3(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_28]), c_0_11])])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk3_2(esk1_0,X1),u1_struct_0(esk2_0))), inference(rw,[status(thm)],[c_0_29, c_0_25])).
cnf(c_0_32, plain, (r2_lattice3(X1,X2,esk3_2(X1,X2))|~v3_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_33, plain, ![X26, X27, X28, X29, X30, X31]:((~r1_orders_2(X26,X28,X29)|r1_orders_2(X27,X30,X31)|(X28!=X30|X29!=X31)|~m1_subset_1(X31,u1_struct_0(X27))|~m1_subset_1(X30,u1_struct_0(X27))|~m1_subset_1(X29,u1_struct_0(X26))|~m1_subset_1(X28,u1_struct_0(X26))|g1_orders_2(u1_struct_0(X26),u1_orders_2(X26))!=g1_orders_2(u1_struct_0(X27),u1_orders_2(X27))|~l1_orders_2(X27)|~l1_orders_2(X26))&(~r2_orders_2(X26,X28,X29)|r2_orders_2(X27,X30,X31)|(X28!=X30|X29!=X31)|~m1_subset_1(X31,u1_struct_0(X27))|~m1_subset_1(X30,u1_struct_0(X27))|~m1_subset_1(X29,u1_struct_0(X26))|~m1_subset_1(X28,u1_struct_0(X26))|g1_orders_2(u1_struct_0(X26),u1_orders_2(X26))!=g1_orders_2(u1_struct_0(X27),u1_orders_2(X27))|~l1_orders_2(X27)|~l1_orders_2(X26))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_yellow_0])])])])).
cnf(c_0_34, negated_conjecture, (r2_lattice3(esk2_0,X1,esk3_2(esk1_0,X2))|~r2_lattice3(esk1_0,X1,esk3_2(esk1_0,X2))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_35, negated_conjecture, (r2_lattice3(esk1_0,X1,esk3_2(esk1_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_27]), c_0_24])])).
cnf(c_0_36, plain, (r1_orders_2(X4,X5,X6)|~r1_orders_2(X1,X2,X3)|X2!=X5|X3!=X6|~m1_subset_1(X6,u1_struct_0(X4))|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))|~l1_orders_2(X4)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_37, negated_conjecture, (r2_lattice3(esk1_0,X1,X2)|g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))!=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))|~r2_lattice3(X3,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(X3))|~l1_orders_2(X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_18]), c_0_24])]), c_0_25]), c_0_25])).
cnf(c_0_38, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|v3_lattice3(X1)|~r2_lattice3(X1,esk4_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_39, negated_conjecture, (r2_lattice3(esk2_0,X1,esk3_2(esk1_0,X1))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40, negated_conjecture, (~v3_lattice3(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_41, plain, (r1_orders_2(X1,X2,X3)|g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_orders_2(X4,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))|~l1_orders_2(X1)|~l1_orders_2(X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_36])])).
cnf(c_0_42, plain, (r1_orders_2(X2,esk3_2(X2,X3),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~v3_lattice3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_43, negated_conjecture, (r2_lattice3(esk1_0,X1,X2)|~r2_lattice3(esk2_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_37]), c_0_11])])).
cnf(c_0_44, negated_conjecture, (m1_subset_1(esk5_2(esk2_0,esk3_2(esk1_0,esk4_1(esk2_0))),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_31]), c_0_11])]), c_0_40])).
cnf(c_0_45, plain, (r2_lattice3(X1,esk4_1(X1),esk5_2(X1,X2))|v3_lattice3(X1)|~r2_lattice3(X1,esk4_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_46, negated_conjecture, (r1_orders_2(X1,X2,X3)|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_orders_2(esk1_0,X2,X3)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_18]), c_0_24])]), c_0_25]), c_0_25]), c_0_25])).
cnf(c_0_47, negated_conjecture, (r1_orders_2(esk1_0,esk3_2(esk1_0,X1),X2)|~r2_lattice3(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_25]), c_0_27]), c_0_24])])).
cnf(c_0_48, negated_conjecture, (r2_lattice3(esk1_0,X1,esk5_2(esk2_0,esk3_2(esk1_0,esk4_1(esk2_0))))|~r2_lattice3(esk2_0,X1,esk5_2(esk2_0,esk3_2(esk1_0,esk4_1(esk2_0))))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_49, negated_conjecture, (r2_lattice3(esk2_0,esk4_1(esk2_0),esk5_2(esk2_0,esk3_2(esk1_0,esk4_1(esk2_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_39]), c_0_31]), c_0_11])]), c_0_40])).
cnf(c_0_50, negated_conjecture, (r1_orders_2(esk2_0,X1,X2)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_46]), c_0_11])])).
cnf(c_0_51, negated_conjecture, (r1_orders_2(esk1_0,esk3_2(esk1_0,X1),esk5_2(esk2_0,esk3_2(esk1_0,esk4_1(esk2_0))))|~r2_lattice3(esk1_0,X1,esk5_2(esk2_0,esk3_2(esk1_0,esk4_1(esk2_0))))), inference(spm,[status(thm)],[c_0_47, c_0_44])).
cnf(c_0_52, negated_conjecture, (r2_lattice3(esk1_0,esk4_1(esk2_0),esk5_2(esk2_0,esk3_2(esk1_0,esk4_1(esk2_0))))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_53, negated_conjecture, (r1_orders_2(esk2_0,X1,esk5_2(esk2_0,esk3_2(esk1_0,esk4_1(esk2_0))))|~r1_orders_2(esk1_0,X1,esk5_2(esk2_0,esk3_2(esk1_0,esk4_1(esk2_0))))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_50, c_0_44])).
cnf(c_0_54, negated_conjecture, (r1_orders_2(esk1_0,esk3_2(esk1_0,esk4_1(esk2_0)),esk5_2(esk2_0,esk3_2(esk1_0,esk4_1(esk2_0))))), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_55, plain, (v3_lattice3(X1)|~r1_orders_2(X1,X2,esk5_2(X1,X2))|~r2_lattice3(X1,esk4_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_56, negated_conjecture, (r1_orders_2(esk2_0,esk3_2(esk1_0,esk4_1(esk2_0)),esk5_2(esk2_0,esk3_2(esk1_0,esk4_1(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_31])])).
cnf(c_0_57, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_39]), c_0_31]), c_0_11])]), c_0_40]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 58
# Proof object clause steps            : 45
# Proof object formula steps           : 13
# Proof object conjectures             : 35
# Proof object clause conjectures      : 32
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 46
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 18
# Processed clauses                    : 86
# ...of these trivial                  : 0
# ...subsumed                          : 6
# ...remaining for further processing  : 80
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 4
# Generated clauses                    : 135
# ...of the previous two non-trivial   : 110
# Contextual simplify-reflections      : 0
# Paramodulations                      : 118
# Factorizations                       : 0
# Equation resolutions                 : 19
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 72
#    Positive orientable unit clauses  : 22
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 49
# Current number of unprocessed clauses: 31
# ...number of literals in the above   : 103
# Current number of archived formulas  : 0
# Current number of archived clauses   : 4
# Clause-clause subsumption calls (NU) : 450
# Rec. Clause-clause subsumption calls : 106
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 17
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5802

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28544 pages
