# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d8_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(X3=k5_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>?[X6]:(r2_hidden(k4_tarski(X4,X6),X1)&r2_hidden(k4_tarski(X6,X5),X2))))))), file('data/MPTP/proofs/05d5901985834853838787a9bc8a1552.E_input', d8_relat_1)).
fof(dt_k5_relat_1, axiom, ![X1, X2]:((v1_relat_1(X1)&v1_relat_1(X2))=>v1_relat_1(k5_relat_1(X1,X2))), file('data/MPTP/proofs/05d5901985834853838787a9bc8a1552.E_input', dt_k5_relat_1)).
fof(d3_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(r1_tarski(X1,X2)<=>![X3, X4]:(r2_hidden(k4_tarski(X3,X4),X1)=>r2_hidden(k4_tarski(X3,X4),X2)))), file('data/MPTP/proofs/05d5901985834853838787a9bc8a1552.E_input', d3_relat_1)).
fof(t49_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>r1_tarski(k5_relat_1(X1,X3),k5_relat_1(X2,X3)))))), file('data/MPTP/proofs/05d5901985834853838787a9bc8a1552.E_input', t49_relat_1)).
fof(c_0_4, plain, ![X10, X11, X12, X13, X14, X16, X17, X18, X21]:((((r2_hidden(k4_tarski(X13,esk4_5(X10,X11,X12,X13,X14)),X10)|~r2_hidden(k4_tarski(X13,X14),X12)|X12!=k5_relat_1(X10,X11)|~v1_relat_1(X12)|~v1_relat_1(X11)|~v1_relat_1(X10))&(r2_hidden(k4_tarski(esk4_5(X10,X11,X12,X13,X14),X14),X11)|~r2_hidden(k4_tarski(X13,X14),X12)|X12!=k5_relat_1(X10,X11)|~v1_relat_1(X12)|~v1_relat_1(X11)|~v1_relat_1(X10)))&(~r2_hidden(k4_tarski(X16,X18),X10)|~r2_hidden(k4_tarski(X18,X17),X11)|r2_hidden(k4_tarski(X16,X17),X12)|X12!=k5_relat_1(X10,X11)|~v1_relat_1(X12)|~v1_relat_1(X11)|~v1_relat_1(X10)))&((~r2_hidden(k4_tarski(esk5_3(X10,X11,X12),esk6_3(X10,X11,X12)),X12)|(~r2_hidden(k4_tarski(esk5_3(X10,X11,X12),X21),X10)|~r2_hidden(k4_tarski(X21,esk6_3(X10,X11,X12)),X11))|X12=k5_relat_1(X10,X11)|~v1_relat_1(X12)|~v1_relat_1(X11)|~v1_relat_1(X10))&((r2_hidden(k4_tarski(esk5_3(X10,X11,X12),esk7_3(X10,X11,X12)),X10)|r2_hidden(k4_tarski(esk5_3(X10,X11,X12),esk6_3(X10,X11,X12)),X12)|X12=k5_relat_1(X10,X11)|~v1_relat_1(X12)|~v1_relat_1(X11)|~v1_relat_1(X10))&(r2_hidden(k4_tarski(esk7_3(X10,X11,X12),esk6_3(X10,X11,X12)),X11)|r2_hidden(k4_tarski(esk5_3(X10,X11,X12),esk6_3(X10,X11,X12)),X12)|X12=k5_relat_1(X10,X11)|~v1_relat_1(X12)|~v1_relat_1(X11)|~v1_relat_1(X10))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_relat_1])])])])])])).
fof(c_0_5, plain, ![X23, X24]:(~v1_relat_1(X23)|~v1_relat_1(X24)|v1_relat_1(k5_relat_1(X23,X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_relat_1])])).
fof(c_0_6, plain, ![X25, X26, X27, X28, X29]:((~r1_tarski(X25,X26)|(~r2_hidden(k4_tarski(X27,X28),X25)|r2_hidden(k4_tarski(X27,X28),X26))|~v1_relat_1(X25))&((r2_hidden(k4_tarski(esk8_2(X25,X29),esk9_2(X25,X29)),X25)|r1_tarski(X25,X29)|~v1_relat_1(X25))&(~r2_hidden(k4_tarski(esk8_2(X25,X29),esk9_2(X25,X29)),X29)|r1_tarski(X25,X29)|~v1_relat_1(X25)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_relat_1])])])])])])).
cnf(c_0_7, plain, (r2_hidden(k4_tarski(X1,esk4_5(X2,X3,X4,X1,X5)),X2)|~r2_hidden(k4_tarski(X1,X5),X4)|X4!=k5_relat_1(X2,X3)|~v1_relat_1(X4)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, plain, (v1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (r2_hidden(k4_tarski(X3,X4),X2)|~r1_tarski(X1,X2)|~r2_hidden(k4_tarski(X3,X4),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (r2_hidden(k4_tarski(X1,esk4_5(X2,X3,k5_relat_1(X2,X3),X1,X4)),X2)|~r2_hidden(k4_tarski(X1,X4),k5_relat_1(X2,X3))|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_7]), c_0_8])).
fof(c_0_11, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>r1_tarski(k5_relat_1(X1,X3),k5_relat_1(X2,X3))))))), inference(assume_negation,[status(cth)],[t49_relat_1])).
cnf(c_0_12, plain, (r2_hidden(k4_tarski(X1,X4),X6)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(k4_tarski(X2,X4),X5)|X6!=k5_relat_1(X3,X5)|~v1_relat_1(X6)|~v1_relat_1(X5)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_13, plain, (r2_hidden(k4_tarski(esk4_5(X1,X2,X3,X4,X5),X5),X2)|~r2_hidden(k4_tarski(X4,X5),X3)|X3!=k5_relat_1(X1,X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_14, plain, (r2_hidden(k4_tarski(X1,esk4_5(X2,X3,k5_relat_1(X2,X3),X1,X4)),X5)|~r2_hidden(k4_tarski(X1,X4),k5_relat_1(X2,X3))|~r1_tarski(X2,X5)|~v1_relat_1(X2)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
fof(c_0_15, negated_conjecture, (v1_relat_1(esk1_0)&(v1_relat_1(esk2_0)&(v1_relat_1(esk3_0)&(r1_tarski(esk1_0,esk2_0)&~r1_tarski(k5_relat_1(esk1_0,esk3_0),k5_relat_1(esk2_0,esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
cnf(c_0_16, plain, (r2_hidden(k4_tarski(X1,X2),k5_relat_1(X3,X4))|~r2_hidden(k4_tarski(X5,X2),X4)|~r2_hidden(k4_tarski(X1,X5),X3)|~v1_relat_1(X4)|~v1_relat_1(X3)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_12]), c_0_8])).
cnf(c_0_17, plain, (r2_hidden(k4_tarski(esk4_5(X1,X2,k5_relat_1(X1,X2),X3,X4),X4),X2)|~r2_hidden(k4_tarski(X3,X4),k5_relat_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_13]), c_0_8])).
cnf(c_0_18, plain, (r2_hidden(k4_tarski(X1,esk4_5(X2,X3,k5_relat_1(X2,X3),X1,X4)),X5)|~r2_hidden(k4_tarski(X1,X4),k5_relat_1(X2,X3))|~r1_tarski(X6,X5)|~r1_tarski(X2,X6)|~v1_relat_1(X6)|~v1_relat_1(X2)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_9, c_0_14])).
cnf(c_0_19, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (r2_hidden(k4_tarski(X1,X2),k5_relat_1(X3,X4))|~r2_hidden(k4_tarski(X1,esk4_5(X5,X4,k5_relat_1(X5,X4),X6,X2)),X3)|~r2_hidden(k4_tarski(X6,X2),k5_relat_1(X5,X4))|~v1_relat_1(X4)|~v1_relat_1(X3)|~v1_relat_1(X5)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_22, negated_conjecture, (r2_hidden(k4_tarski(X1,esk4_5(X2,X3,k5_relat_1(X2,X3),X1,X4)),esk2_0)|~r2_hidden(k4_tarski(X1,X4),k5_relat_1(X2,X3))|~r1_tarski(X2,esk1_0)|~v1_relat_1(X2)|~v1_relat_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])])).
cnf(c_0_23, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, negated_conjecture, (r2_hidden(k4_tarski(X1,X2),k5_relat_1(esk2_0,X3))|~r2_hidden(k4_tarski(X1,X2),k5_relat_1(X4,X3))|~r1_tarski(X4,esk1_0)|~v1_relat_1(X3)|~v1_relat_1(X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])])).
cnf(c_0_25, plain, (r2_hidden(k4_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),X1)|r1_tarski(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_26, plain, (r1_tarski(X1,X2)|~r2_hidden(k4_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_27, negated_conjecture, (r2_hidden(k4_tarski(esk8_2(k5_relat_1(X1,X2),X3),esk9_2(k5_relat_1(X1,X2),X3)),k5_relat_1(esk2_0,X2))|r1_tarski(k5_relat_1(X1,X2),X3)|~r1_tarski(X1,esk1_0)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_8])).
cnf(c_0_28, negated_conjecture, (~r1_tarski(k5_relat_1(esk1_0,esk3_0),k5_relat_1(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_29, negated_conjecture, (r1_tarski(k5_relat_1(X1,X2),k5_relat_1(esk2_0,X2))|~r1_tarski(X1,esk1_0)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_8])).
cnf(c_0_30, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_31, negated_conjecture, (~r1_tarski(esk1_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30]), c_0_20])])).
cnf(c_0_32, plain, (r1_tarski(X1,X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_26, c_0_25])).
cnf(c_0_33, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_20])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 34
# Proof object clause steps            : 25
# Proof object formula steps           : 9
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 15
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 15
# Processed clauses                    : 174
# ...of these trivial                  : 0
# ...subsumed                          : 28
# ...remaining for further processing  : 146
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 554
# ...of the previous two non-trivial   : 550
# Contextual simplify-reflections      : 12
# Paramodulations                      : 551
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 128
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 122
# Current number of unprocessed clauses: 406
# ...number of literals in the above   : 3868
# Current number of archived formulas  : 0
# Current number of archived clauses   : 15
# Clause-clause subsumption calls (NU) : 14388
# Rec. Clause-clause subsumption calls : 3164
# Non-unit clause-clause subsumptions  : 40
# Unit Clause-clause subsumption calls : 121
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 21134

# -------------------------------------------------
# User time                : 0.027 s
# System time              : 0.000 s
# Total time               : 0.027 s
# Maximum resident set size: 28808 pages
