# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t166_funct_1, conjecture, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((k1_relat_1(X1)=k1_relat_1(X2)&![X4]:(r2_hidden(X4,X3)=>k1_funct_1(X1,X4)=k1_funct_1(X2,X4)))=>k7_relat_1(X1,X3)=k7_relat_1(X2,X3)))), file('data/MPTP/proofs/05eaa70520b04db8b27c15482b042ada.E_input', t166_funct_1)).
fof(t91_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(X1,k1_relat_1(X2))=>k1_relat_1(k7_relat_1(X2,X1))=X1)), file('data/MPTP/proofs/05eaa70520b04db8b27c15482b042ada.E_input', t91_relat_1)).
fof(t87_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k1_relat_1(k7_relat_1(X2,X1)),X1)), file('data/MPTP/proofs/05eaa70520b04db8b27c15482b042ada.E_input', t87_relat_1)).
fof(t189_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>k7_relat_1(X1,k1_relat_1(X2))=k7_relat_1(X1,k1_relat_1(k7_relat_1(X2,k1_relat_1(X1)))))), file('data/MPTP/proofs/05eaa70520b04db8b27c15482b042ada.E_input', t189_relat_1)).
fof(t89_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k1_relat_1(k7_relat_1(X2,X1)),k1_relat_1(X2))), file('data/MPTP/proofs/05eaa70520b04db8b27c15482b042ada.E_input', t89_relat_1)).
fof(s3_funct_1__e2_24__funct_1, axiom, ![X1, X2]:?[X3]:((v1_relat_1(X3)&v1_funct_1(X3))&(k1_relat_1(X3)=X1&![X4]:(r2_hidden(X4,X1)=>k1_funct_1(X3,X4)=X2))), file('data/MPTP/proofs/05eaa70520b04db8b27c15482b042ada.E_input', s3_funct_1__e2_24__funct_1)).
fof(t165_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((r1_tarski(X3,k1_relat_1(X1))&r1_tarski(X3,k1_relat_1(X2)))=>(k7_relat_1(X1,X3)=k7_relat_1(X2,X3)<=>![X4]:(r2_hidden(X4,X3)=>k1_funct_1(X1,X4)=k1_funct_1(X2,X4)))))), file('data/MPTP/proofs/05eaa70520b04db8b27c15482b042ada.E_input', t165_funct_1)).
fof(t86_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k1_relat_1(k7_relat_1(X3,X2)))<=>(r2_hidden(X1,X2)&r2_hidden(X1,k1_relat_1(X3))))), file('data/MPTP/proofs/05eaa70520b04db8b27c15482b042ada.E_input', t86_relat_1)).
fof(t97_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(k1_relat_1(X2),X1)=>k7_relat_1(X2,X1)=X2)), file('data/MPTP/proofs/05eaa70520b04db8b27c15482b042ada.E_input', t97_relat_1)).
fof(c_0_9, negated_conjecture, ~(![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((k1_relat_1(X1)=k1_relat_1(X2)&![X4]:(r2_hidden(X4,X3)=>k1_funct_1(X1,X4)=k1_funct_1(X2,X4)))=>k7_relat_1(X1,X3)=k7_relat_1(X2,X3))))), inference(assume_negation,[status(cth)],[t166_funct_1])).
fof(c_0_10, plain, ![X9, X10]:(~v1_relat_1(X10)|(~r1_tarski(X9,k1_relat_1(X10))|k1_relat_1(k7_relat_1(X10,X9))=X9)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t91_relat_1])])).
fof(c_0_11, negated_conjecture, ![X8]:((v1_relat_1(esk1_0)&v1_funct_1(esk1_0))&((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&((k1_relat_1(esk1_0)=k1_relat_1(esk2_0)&(~r2_hidden(X8,esk3_0)|k1_funct_1(esk1_0,X8)=k1_funct_1(esk2_0,X8)))&k7_relat_1(esk1_0,esk3_0)!=k7_relat_1(esk2_0,esk3_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])).
cnf(c_0_12, plain, (k1_relat_1(k7_relat_1(X1,X2))=X2|~v1_relat_1(X1)|~r1_tarski(X2,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_13, negated_conjecture, (k1_relat_1(esk1_0)=k1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_14, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_15, plain, ![X11, X12]:(~v1_relat_1(X12)|r1_tarski(k1_relat_1(k7_relat_1(X12,X11)),X11)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t87_relat_1])])).
fof(c_0_16, plain, ![X13, X14]:(~v1_relat_1(X13)|(~v1_relat_1(X14)|k7_relat_1(X13,k1_relat_1(X14))=k7_relat_1(X13,k1_relat_1(k7_relat_1(X14,k1_relat_1(X13)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t189_relat_1])])])).
cnf(c_0_17, negated_conjecture, (k1_relat_1(k7_relat_1(esk1_0,X1))=X1|~r1_tarski(X1,k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14])])).
cnf(c_0_18, plain, (r1_tarski(k1_relat_1(k7_relat_1(X1,X2)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, plain, (k7_relat_1(X1,k1_relat_1(X2))=k7_relat_1(X1,k1_relat_1(k7_relat_1(X2,k1_relat_1(X1))))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_20, plain, ![X15, X16]:(~v1_relat_1(X16)|r1_tarski(k1_relat_1(k7_relat_1(X16,X15)),k1_relat_1(X16))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t89_relat_1])])).
cnf(c_0_21, negated_conjecture, (k1_relat_1(k7_relat_1(esk1_0,k1_relat_1(k7_relat_1(X1,k1_relat_1(esk2_0)))))=k1_relat_1(k7_relat_1(X1,k1_relat_1(esk2_0)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_22, negated_conjecture, (k7_relat_1(esk1_0,k1_relat_1(k7_relat_1(X1,k1_relat_1(esk2_0))))=k7_relat_1(esk1_0,k1_relat_1(X1))|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_13]), c_0_14])])).
cnf(c_0_23, plain, (r1_tarski(k1_relat_1(k7_relat_1(X1,X2)),k1_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24, negated_conjecture, (k1_relat_1(k7_relat_1(esk1_0,k1_relat_1(X1)))=k1_relat_1(k7_relat_1(X1,k1_relat_1(esk2_0)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_25, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_26, plain, ![X17, X18, X20]:((v1_relat_1(esk4_2(X17,X18))&v1_funct_1(esk4_2(X17,X18)))&(k1_relat_1(esk4_2(X17,X18))=X17&(~r2_hidden(X20,X17)|k1_funct_1(esk4_2(X17,X18),X20)=X18))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[s3_funct_1__e2_24__funct_1])])])])).
cnf(c_0_27, negated_conjecture, (r1_tarski(k1_relat_1(k7_relat_1(esk1_0,X1)),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_13]), c_0_14])])).
cnf(c_0_28, negated_conjecture, (k7_relat_1(esk2_0,k1_relat_1(k7_relat_1(esk1_0,k1_relat_1(X1))))=k7_relat_1(esk2_0,k1_relat_1(X1))|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_24]), c_0_25])])).
cnf(c_0_29, plain, (k1_relat_1(esk4_2(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_30, plain, (v1_relat_1(esk4_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_31, plain, ![X21, X22, X23, X24]:((k7_relat_1(X21,X23)!=k7_relat_1(X22,X23)|(~r2_hidden(X24,X23)|k1_funct_1(X21,X24)=k1_funct_1(X22,X24))|(~r1_tarski(X23,k1_relat_1(X21))|~r1_tarski(X23,k1_relat_1(X22)))|(~v1_relat_1(X22)|~v1_funct_1(X22))|(~v1_relat_1(X21)|~v1_funct_1(X21)))&((r2_hidden(esk5_3(X21,X22,X23),X23)|k7_relat_1(X21,X23)=k7_relat_1(X22,X23)|(~r1_tarski(X23,k1_relat_1(X21))|~r1_tarski(X23,k1_relat_1(X22)))|(~v1_relat_1(X22)|~v1_funct_1(X22))|(~v1_relat_1(X21)|~v1_funct_1(X21)))&(k1_funct_1(X21,esk5_3(X21,X22,X23))!=k1_funct_1(X22,esk5_3(X21,X22,X23))|k7_relat_1(X21,X23)=k7_relat_1(X22,X23)|(~r1_tarski(X23,k1_relat_1(X21))|~r1_tarski(X23,k1_relat_1(X22)))|(~v1_relat_1(X22)|~v1_funct_1(X22))|(~v1_relat_1(X21)|~v1_funct_1(X21))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t165_funct_1])])])])])).
fof(c_0_32, plain, ![X26, X27, X28]:(((r2_hidden(X26,X27)|~r2_hidden(X26,k1_relat_1(k7_relat_1(X28,X27)))|~v1_relat_1(X28))&(r2_hidden(X26,k1_relat_1(X28))|~r2_hidden(X26,k1_relat_1(k7_relat_1(X28,X27)))|~v1_relat_1(X28)))&(~r2_hidden(X26,X27)|~r2_hidden(X26,k1_relat_1(X28))|r2_hidden(X26,k1_relat_1(k7_relat_1(X28,X27)))|~v1_relat_1(X28))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t86_relat_1])])])).
fof(c_0_33, plain, ![X29, X30]:(~v1_relat_1(X30)|(~r1_tarski(k1_relat_1(X30),X29)|k7_relat_1(X30,X29)=X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t97_relat_1])])).
cnf(c_0_34, negated_conjecture, (k1_relat_1(k7_relat_1(esk2_0,k1_relat_1(k7_relat_1(esk1_0,X1))))=k1_relat_1(k7_relat_1(esk1_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_27]), c_0_25])])).
cnf(c_0_35, negated_conjecture, (k7_relat_1(esk2_0,k1_relat_1(k7_relat_1(esk1_0,X1)))=k7_relat_1(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])])).
cnf(c_0_36, plain, (k7_relat_1(X1,X3)=k7_relat_1(X2,X3)|k1_funct_1(X1,esk5_3(X1,X2,X3))!=k1_funct_1(X2,esk5_3(X1,X2,X3))|~r1_tarski(X3,k1_relat_1(X1))|~r1_tarski(X3,k1_relat_1(X2))|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_37, negated_conjecture, (k1_funct_1(esk1_0,X1)=k1_funct_1(esk2_0,X1)|~r2_hidden(X1,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_38, negated_conjecture, (v1_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_39, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_relat_1(k7_relat_1(X3,X2)))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40, plain, (r2_hidden(esk5_3(X1,X2,X3),X3)|k7_relat_1(X1,X3)=k7_relat_1(X2,X3)|~r1_tarski(X3,k1_relat_1(X1))|~r1_tarski(X3,k1_relat_1(X2))|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41, plain, (k7_relat_1(X1,X2)=X1|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, negated_conjecture, (k1_relat_1(k7_relat_1(esk1_0,X1))=k1_relat_1(k7_relat_1(esk2_0,X1))), inference(rw,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_43, negated_conjecture, (k7_relat_1(X1,X2)=k7_relat_1(esk1_0,X2)|k1_funct_1(X1,esk5_3(X1,esk1_0,X2))!=k1_funct_1(esk2_0,esk5_3(X1,esk1_0,X2))|~r1_tarski(X2,k1_relat_1(esk2_0))|~r1_tarski(X2,k1_relat_1(X1))|~r2_hidden(esk5_3(X1,esk1_0,X2),esk3_0)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_13]), c_0_38]), c_0_14])])).
cnf(c_0_44, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_45, plain, (k7_relat_1(X1,k1_relat_1(k7_relat_1(X2,X3)))=k7_relat_1(X4,k1_relat_1(k7_relat_1(X2,X3)))|r2_hidden(esk5_3(X1,X4,k1_relat_1(k7_relat_1(X2,X3))),X3)|~r1_tarski(k1_relat_1(k7_relat_1(X2,X3)),k1_relat_1(X4))|~r1_tarski(k1_relat_1(k7_relat_1(X2,X3)),k1_relat_1(X1))|~v1_funct_1(X4)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X4)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_46, plain, (k7_relat_1(esk4_2(X1,X2),X3)=esk4_2(X1,X2)|~r1_tarski(X1,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_29]), c_0_30])])).
cnf(c_0_47, negated_conjecture, (k1_relat_1(k7_relat_1(esk2_0,k1_relat_1(X1)))=k1_relat_1(k7_relat_1(X1,k1_relat_1(esk2_0)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_24, c_0_42])).
cnf(c_0_48, negated_conjecture, (k7_relat_1(esk1_0,X1)=k7_relat_1(esk2_0,X1)|~r1_tarski(X1,k1_relat_1(esk2_0))|~r2_hidden(esk5_3(esk2_0,esk1_0,X1),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_43]), c_0_44]), c_0_25])])).
cnf(c_0_49, plain, (k7_relat_1(X1,X2)=k7_relat_1(X3,X2)|r2_hidden(esk5_3(X1,X3,X2),X4)|~r1_tarski(X2,k1_relat_1(X3))|~r1_tarski(X2,k1_relat_1(X1))|~r1_tarski(X2,X4)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_29]), c_0_29]), c_0_29]), c_0_29]), c_0_29]), c_0_30])])).
cnf(c_0_50, negated_conjecture, (k7_relat_1(esk1_0,k1_relat_1(k7_relat_1(esk2_0,k1_relat_1(X1))))=k7_relat_1(esk1_0,k1_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_22, c_0_47])).
cnf(c_0_51, negated_conjecture, (k7_relat_1(esk1_0,X1)=k7_relat_1(esk2_0,X1)|~r1_tarski(X1,k1_relat_1(esk2_0))|~r1_tarski(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_13]), c_0_38]), c_0_44]), c_0_14]), c_0_25])])).
cnf(c_0_52, negated_conjecture, (k7_relat_1(esk1_0,k1_relat_1(k7_relat_1(esk2_0,X1)))=k7_relat_1(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_29]), c_0_30])])).
cnf(c_0_53, negated_conjecture, (k7_relat_1(esk2_0,k1_relat_1(k7_relat_1(esk2_0,X1)))=k7_relat_1(esk2_0,X1)), inference(rw,[status(thm)],[c_0_35, c_0_42])).
cnf(c_0_54, negated_conjecture, (k7_relat_1(esk1_0,X1)=k7_relat_1(esk2_0,X1)|~r1_tarski(k1_relat_1(k7_relat_1(esk2_0,X1)),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_23]), c_0_52]), c_0_53]), c_0_25])])).
cnf(c_0_55, negated_conjecture, (k7_relat_1(esk1_0,esk3_0)!=k7_relat_1(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_56, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_18]), c_0_25])]), c_0_55]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 57
# Proof object clause steps            : 38
# Proof object formula steps           : 19
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 46
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 22
# Processed clauses                    : 3662
# ...of these trivial                  : 18
# ...subsumed                          : 3138
# ...remaining for further processing  : 506
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 20
# Backward-rewritten                   : 29
# Generated clauses                    : 69300
# ...of the previous two non-trivial   : 65798
# Contextual simplify-reflections      : 5
# Paramodulations                      : 69289
# Factorizations                       : 0
# Equation resolutions                 : 11
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 435
#    Positive orientable unit clauses  : 22
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 411
# Current number of unprocessed clauses: 61802
# ...number of literals in the above   : 352876
# Current number of archived formulas  : 0
# Current number of archived clauses   : 71
# Clause-clause subsumption calls (NU) : 95569
# Rec. Clause-clause subsumption calls : 53157
# Non-unit clause-clause subsumptions  : 3162
# Unit Clause-clause subsumption calls : 574
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 99
# BW rewrite match successes           : 18
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2965530

# -------------------------------------------------
# User time                : 0.977 s
# System time              : 0.033 s
# Total time               : 1.010 s
# Maximum resident set size: 58208 pages
