# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/05ef149eb68945c7944b93ff1a02fe39.E_input', d1_xboole_0)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/05ef149eb68945c7944b93ff1a02fe39.E_input', t7_xboole_0)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/05ef149eb68945c7944b93ff1a02fe39.E_input', t5_subset)).
fof(t11_relset_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>((r1_tarski(k1_relat_1(X3),X1)&r1_tarski(k2_relat_1(X3),X2))=>m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))), file('data/MPTP/proofs/05ef149eb68945c7944b93ff1a02fe39.E_input', t11_relset_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/05ef149eb68945c7944b93ff1a02fe39.E_input', d10_xboole_0)).
fof(fc10_relat_1, axiom, ![X1]:(v1_xboole_0(X1)=>v1_xboole_0(k1_relat_1(X1))), file('data/MPTP/proofs/05ef149eb68945c7944b93ff1a02fe39.E_input', fc10_relat_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/05ef149eb68945c7944b93ff1a02fe39.E_input', t113_zfmisc_1)).
fof(t65_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>(k1_relat_1(X1)=k1_xboole_0<=>k2_relat_1(X1)=k1_xboole_0)), file('data/MPTP/proofs/05ef149eb68945c7944b93ff1a02fe39.E_input', t65_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/05ef149eb68945c7944b93ff1a02fe39.E_input', fc6_relat_1)).
fof(t8_funct_2, conjecture, ![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(r1_tarski(k2_relset_1(X1,X2,X4),X3)=>((X2=k1_xboole_0&~(X1=k1_xboole_0))|(v1_funct_1(X4)&(v1_funct_2(X4,X1,X3)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))))))), file('data/MPTP/proofs/05ef149eb68945c7944b93ff1a02fe39.E_input', t8_funct_2)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/05ef149eb68945c7944b93ff1a02fe39.E_input', redefinition_k2_relset_1)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(((X2=k1_xboole_0=>X1=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X2,X3)))&(X2=k1_xboole_0=>(X1=k1_xboole_0|(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0))))), file('data/MPTP/proofs/05ef149eb68945c7944b93ff1a02fe39.E_input', d1_funct_2)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/05ef149eb68945c7944b93ff1a02fe39.E_input', redefinition_k1_relset_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/05ef149eb68945c7944b93ff1a02fe39.E_input', fc1_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/05ef149eb68945c7944b93ff1a02fe39.E_input', d3_tarski)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/05ef149eb68945c7944b93ff1a02fe39.E_input', cc2_relat_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/05ef149eb68945c7944b93ff1a02fe39.E_input', t3_subset)).
fof(c_0_17, plain, ![X5, X6, X7]:((~v1_xboole_0(X5)|~r2_hidden(X6,X5))&(r2_hidden(esk1_1(X7),X7)|v1_xboole_0(X7))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_18, plain, ![X9]:(X9=k1_xboole_0|r2_hidden(esk2_1(X9),X9)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
fof(c_0_19, plain, ![X11, X12, X13]:(~r2_hidden(X11,X12)|(~m1_subset_1(X12,k1_zfmisc_1(X13))|~v1_xboole_0(X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_20, plain, ![X14, X15, X16]:(~v1_relat_1(X16)|(~r1_tarski(k1_relat_1(X16),X14)|~r1_tarski(k2_relat_1(X16),X15)|m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_relset_1])])).
fof(c_0_21, plain, ![X20, X21]:(((r1_tarski(X20,X21)|X20!=X21)&(r1_tarski(X21,X20)|X20!=X21))&(~r1_tarski(X20,X21)|~r1_tarski(X21,X20)|X20=X21)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_22, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (X1=k1_xboole_0|r2_hidden(esk2_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_24, plain, ![X19]:(~v1_xboole_0(X19)|v1_xboole_0(k1_relat_1(X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc10_relat_1])])).
fof(c_0_25, plain, ![X17, X18]:((k2_zfmisc_1(X17,X18)!=k1_xboole_0|(X17=k1_xboole_0|X18=k1_xboole_0))&((X17!=k1_xboole_0|k2_zfmisc_1(X17,X18)=k1_xboole_0)&(X18!=k1_xboole_0|k2_zfmisc_1(X17,X18)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_26, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_29, plain, ![X22]:((k1_relat_1(X22)!=k1_xboole_0|k2_relat_1(X22)=k1_xboole_0|~v1_relat_1(X22))&(k2_relat_1(X22)!=k1_xboole_0|k1_relat_1(X22)=k1_xboole_0|~v1_relat_1(X22))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t65_relat_1])])])).
cnf(c_0_30, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_31, plain, (v1_xboole_0(k1_relat_1(X1))|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_32, plain, ![X23, X24]:v1_relat_1(k2_zfmisc_1(X23,X24)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_33, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_34, negated_conjecture, ~(![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(r1_tarski(k2_relset_1(X1,X2,X4),X3)=>((X2=k1_xboole_0&~(X1=k1_xboole_0))|(v1_funct_1(X4)&(v1_funct_2(X4,X1,X3)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))))))), inference(assume_negation,[status(cth)],[t8_funct_2])).
cnf(c_0_35, plain, (~r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X3)|~v1_relat_1(X1)|~r2_hidden(X4,X1)|~v1_xboole_0(k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_36, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_28])).
cnf(c_0_37, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_38, plain, (k2_relat_1(X1)=k1_xboole_0|k1_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, plain, (k1_relat_1(X1)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_40, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_41, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_33])).
fof(c_0_42, plain, ![X35, X36, X37]:(~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(X35,X36)))|k2_relset_1(X35,X36,X37)=k2_relat_1(X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
fof(c_0_43, negated_conjecture, ((v1_funct_1(esk6_0)&(v1_funct_2(esk6_0,esk3_0,esk4_0)&m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk4_0)))))&(r1_tarski(k2_relset_1(esk3_0,esk4_0,esk6_0),esk5_0)&((esk4_0!=k1_xboole_0|esk3_0=k1_xboole_0)&(~v1_funct_1(esk6_0)|(~v1_funct_2(esk6_0,esk3_0,esk5_0)|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk5_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_34])])])])).
fof(c_0_44, plain, ![X43, X44, X45]:((((~v1_funct_2(X45,X43,X44)|X43=k1_relset_1(X43,X44,X45)|X44=k1_xboole_0|~m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(X43,X44))))&(X43!=k1_relset_1(X43,X44,X45)|v1_funct_2(X45,X43,X44)|X44=k1_xboole_0|~m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(X43,X44)))))&((~v1_funct_2(X45,X43,X44)|X43=k1_relset_1(X43,X44,X45)|X43!=k1_xboole_0|~m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(X43,X44))))&(X43!=k1_relset_1(X43,X44,X45)|v1_funct_2(X45,X43,X44)|X43!=k1_xboole_0|~m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(X43,X44))))))&((~v1_funct_2(X45,X43,X44)|X45=k1_xboole_0|X43=k1_xboole_0|X44!=k1_xboole_0|~m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(X43,X44))))&(X45!=k1_xboole_0|v1_funct_2(X45,X43,X44)|X43=k1_xboole_0|X44!=k1_xboole_0|~m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(X43,X44)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
fof(c_0_45, plain, ![X38, X39, X40]:(~m1_subset_1(X40,k1_zfmisc_1(k2_zfmisc_1(X38,X39)))|k1_relset_1(X38,X39,X40)=k1_relat_1(X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_46, plain, (~r1_tarski(k2_relat_1(X1),X2)|~v1_relat_1(X1)|~r2_hidden(X3,X1)|~v1_xboole_0(k2_zfmisc_1(k1_relat_1(X1),X2))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_47, plain, (k1_relat_1(X1)=k1_xboole_0|k2_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_48, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_49, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_50, plain, (k2_relat_1(X1)=k1_xboole_0|~v1_relat_1(X1)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_51, plain, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
fof(c_0_52, plain, ![X29, X30, X31, X32, X33]:((~r1_tarski(X29,X30)|(~r2_hidden(X31,X29)|r2_hidden(X31,X30)))&((r2_hidden(esk7_2(X32,X33),X32)|r1_tarski(X32,X33))&(~r2_hidden(esk7_2(X32,X33),X33)|r1_tarski(X32,X33)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_53, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk4_0)))), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_55, plain, ![X46, X47]:(~v1_relat_1(X46)|(~m1_subset_1(X47,k1_zfmisc_1(X46))|v1_relat_1(X47))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
cnf(c_0_56, plain, (v1_funct_2(X3,X1,X2)|X2=k1_xboole_0|X1!=k1_relset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_57, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_58, plain, (X2=k1_relset_1(X2,X3,X1)|X3=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_59, negated_conjecture, (v1_funct_2(esk6_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_60, plain, (k2_relat_1(X1)!=k1_xboole_0|~r1_tarski(k2_relat_1(X1),X2)|~v1_relat_1(X1)|~r2_hidden(X3,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48]), c_0_49])])).
cnf(c_0_61, plain, (k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_49])])).
cnf(c_0_62, plain, (r2_hidden(esk7_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
fof(c_0_63, plain, ![X41, X42]:((~m1_subset_1(X41,k1_zfmisc_1(X42))|r1_tarski(X41,X42))&(~r1_tarski(X41,X42)|m1_subset_1(X41,k1_zfmisc_1(X42)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_64, negated_conjecture, (~v1_funct_1(esk6_0)|~v1_funct_2(esk6_0,esk3_0,esk5_0)|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk5_0)))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_65, negated_conjecture, (v1_funct_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_66, negated_conjecture, (r1_tarski(k2_relset_1(esk3_0,esk4_0,esk6_0),esk5_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_67, negated_conjecture, (k2_relset_1(esk3_0,esk4_0,esk6_0)=k2_relat_1(esk6_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_68, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_69, plain, (X1=k1_xboole_0|v1_funct_2(X2,X3,X1)|k1_relset_1(X3,X1,X2)!=X3|~r1_tarski(k1_relat_1(X2),X3)|~r1_tarski(k2_relat_1(X2),X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_56, c_0_27])).
cnf(c_0_70, plain, (k1_relset_1(X1,X2,X3)=k1_relat_1(X3)|~r1_tarski(k1_relat_1(X3),X1)|~r1_tarski(k2_relat_1(X3),X2)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_57, c_0_27])).
cnf(c_0_71, negated_conjecture, (k1_relset_1(esk3_0,esk4_0,esk6_0)=esk3_0|k1_xboole_0=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_54]), c_0_59])])).
cnf(c_0_72, negated_conjecture, (k1_relset_1(esk3_0,esk4_0,esk6_0)=k1_relat_1(esk6_0)), inference(spm,[status(thm)],[c_0_57, c_0_54])).
cnf(c_0_73, plain, (~r1_tarski(k1_xboole_0,X1)|~r2_hidden(X2,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_51])])).
cnf(c_0_74, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_22, c_0_62])).
cnf(c_0_75, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_76, negated_conjecture, (~v1_funct_2(esk6_0,esk3_0,esk5_0)|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk5_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_65])])).
cnf(c_0_77, negated_conjecture, (r1_tarski(k2_relat_1(esk6_0),esk5_0)), inference(rw,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_78, negated_conjecture, (v1_relat_1(esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_54]), c_0_40])])).
cnf(c_0_79, plain, (X1=k1_xboole_0|v1_funct_2(X2,k1_relat_1(X2),X1)|~r1_tarski(k2_relat_1(X2),X1)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70])]), c_0_36])])).
cnf(c_0_80, negated_conjecture, (k1_relat_1(esk6_0)=esk3_0|k1_xboole_0=esk4_0), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_81, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_49])])).
cnf(c_0_82, plain, (~r1_tarski(X1,X2)|~r2_hidden(X3,X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_26, c_0_75])).
cnf(c_0_83, negated_conjecture, (~v1_funct_2(esk6_0,esk3_0,esk5_0)|~r1_tarski(k1_relat_1(esk6_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_27]), c_0_77]), c_0_78])])).
cnf(c_0_84, negated_conjecture, (k1_xboole_0=esk4_0|X1=k1_xboole_0|v1_funct_2(esk6_0,esk3_0,X1)|~r1_tarski(k2_relat_1(esk6_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_78])])).
cnf(c_0_85, plain, (k1_relset_1(X1,X2,X3)=k1_relat_1(X3)|~r1_tarski(X3,k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_57, c_0_75])).
cnf(c_0_86, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_81, c_0_62])).
cnf(c_0_87, plain, (~r1_tarski(k2_relat_1(X1),k1_xboole_0)|~v1_relat_1(X1)|~r2_hidden(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_41]), c_0_49])])).
cnf(c_0_88, plain, (r2_hidden(esk1_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_89, negated_conjecture, (~r2_hidden(X1,k2_relat_1(esk6_0))|~v1_xboole_0(esk5_0)), inference(spm,[status(thm)],[c_0_82, c_0_77])).
cnf(c_0_90, negated_conjecture, (k1_xboole_0=esk5_0|k1_xboole_0=esk4_0|~r1_tarski(k1_relat_1(esk6_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_77])])).
cnf(c_0_91, plain, (X2=k1_relset_1(X2,X3,X1)|~v1_funct_2(X1,X2,X3)|X2!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_92, plain, (k1_relset_1(X1,X2,k1_xboole_0)=k1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_93, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_94, plain, (v1_xboole_0(X1)|~r1_tarski(k2_relat_1(X1),k1_xboole_0)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_95, negated_conjecture, (r1_tarski(k2_relat_1(esk6_0),X1)|~v1_xboole_0(esk5_0)), inference(spm,[status(thm)],[c_0_89, c_0_62])).
cnf(c_0_96, negated_conjecture, (k1_xboole_0=esk4_0|k1_xboole_0=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_80]), c_0_36])])).
cnf(c_0_97, plain, (k1_relset_1(k1_xboole_0,X1,X2)=k1_xboole_0|~v1_funct_2(X2,k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0))), inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_91]), c_0_48])).
cnf(c_0_98, plain, (X1=k1_xboole_0|v1_funct_2(k1_xboole_0,k1_relat_1(k1_xboole_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_92]), c_0_61]), c_0_86]), c_0_51])])]), c_0_36])])).
cnf(c_0_99, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_100, negated_conjecture, (r1_tarski(esk6_0,k2_zfmisc_1(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_93, c_0_54])).
cnf(c_0_101, negated_conjecture, (esk3_0=k1_xboole_0|esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_102, negated_conjecture, (v1_xboole_0(esk6_0)|~v1_xboole_0(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_78])])).
cnf(c_0_103, negated_conjecture, (k1_xboole_0=esk4_0|v1_xboole_0(esk5_0)), inference(spm,[status(thm)],[c_0_49, c_0_96])).
cnf(c_0_104, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0|~v1_funct_2(k1_xboole_0,k1_xboole_0,X1)|~m1_subset_1(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_97, c_0_92])).
cnf(c_0_105, plain, (X1=k1_xboole_0|v1_funct_2(k1_xboole_0,k1_xboole_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_47]), c_0_61]), c_0_51])])).
cnf(c_0_106, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(esk3_0,esk4_0))|~r2_hidden(X1,esk6_0)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_107, negated_conjecture, (k1_xboole_0=esk3_0|~v1_xboole_0(esk6_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_80]), c_0_101])).
cnf(c_0_108, negated_conjecture, (k1_xboole_0=esk4_0|v1_xboole_0(esk6_0)), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_109, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_110, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_111, plain, (r1_tarski(X1,X2)|~r1_tarski(k2_relat_1(X1),k1_xboole_0)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_87, c_0_62])).
cnf(c_0_112, negated_conjecture, (~r1_tarski(k2_relat_1(k2_zfmisc_1(esk3_0,esk4_0)),k1_xboole_0)|~r2_hidden(X1,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_106]), c_0_40])])).
cnf(c_0_113, negated_conjecture, (k1_xboole_0=esk3_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_101])).
cnf(c_0_114, plain, (v1_funct_2(X3,X1,X2)|X1!=k1_relset_1(X1,X2,X3)|X1!=k1_xboole_0|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_115, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0|X1=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_75]), c_0_36])])).
cnf(c_0_116, negated_conjecture, (k2_zfmisc_1(esk3_0,esk4_0)=esk6_0|~r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),esk6_0)), inference(spm,[status(thm)],[c_0_110, c_0_100])).
cnf(c_0_117, plain, (r1_tarski(X1,X2)|~v1_relat_1(X1)|~v1_xboole_0(k2_relat_1(X1))), inference(spm,[status(thm)],[c_0_111, c_0_74])).
cnf(c_0_118, negated_conjecture, (~r2_hidden(X1,esk6_0)|~v1_xboole_0(k2_relat_1(k2_zfmisc_1(esk3_0,esk4_0)))), inference(spm,[status(thm)],[c_0_112, c_0_74])).
cnf(c_0_119, plain, (k2_zfmisc_1(esk3_0,X1)=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_113]), c_0_113])).
cnf(c_0_120, plain, (k2_relat_1(esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_113]), c_0_113])).
cnf(c_0_121, plain, (v1_xboole_0(esk3_0)), inference(rw,[status(thm)],[c_0_49, c_0_113])).
cnf(c_0_122, plain, (v1_funct_2(X1,k1_xboole_0,X2)|k1_relset_1(k1_xboole_0,X2,X1)!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_114]), c_0_48])).
cnf(c_0_123, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~r1_tarski(k2_relat_1(X1),k1_xboole_0)|~r1_tarski(k1_relat_1(X1),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_27, c_0_41])).
cnf(c_0_124, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(ef,[status(thm)],[c_0_115])).
cnf(c_0_125, negated_conjecture, (~v1_funct_2(esk6_0,esk3_0,esk5_0)|~r1_tarski(esk6_0,k2_zfmisc_1(esk3_0,esk5_0))), inference(spm,[status(thm)],[c_0_76, c_0_75])).
cnf(c_0_126, negated_conjecture, (k2_zfmisc_1(esk3_0,esk4_0)=esk6_0|~v1_xboole_0(k2_relat_1(k2_zfmisc_1(esk3_0,esk4_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_40])])).
cnf(c_0_127, negated_conjecture, (~r2_hidden(X1,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_118, c_0_119]), c_0_120]), c_0_121])])).
cnf(c_0_128, plain, (v1_funct_2(k1_xboole_0,k1_xboole_0,X1)|k1_relat_1(k1_xboole_0)!=k1_xboole_0|~m1_subset_1(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_122, c_0_92])).
cnf(c_0_129, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124]), c_0_61]), c_0_36]), c_0_86]), c_0_51])])).
cnf(c_0_130, negated_conjecture, (~v1_funct_2(esk6_0,esk3_0,esk5_0)|~v1_xboole_0(esk6_0)), inference(spm,[status(thm)],[c_0_125, c_0_74])).
cnf(c_0_131, negated_conjecture, (esk3_0=esk6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_126, c_0_119]), c_0_119]), c_0_120]), c_0_121])])).
cnf(c_0_132, negated_conjecture, (v1_xboole_0(esk6_0)), inference(spm,[status(thm)],[c_0_127, c_0_88])).
cnf(c_0_133, plain, (v1_funct_2(k1_xboole_0,k1_xboole_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_128, c_0_124])]), c_0_129])])).
cnf(c_0_134, negated_conjecture, (~v1_funct_2(esk6_0,esk6_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_130, c_0_131]), c_0_132])])).
cnf(c_0_135, plain, (v1_funct_2(esk6_0,esk6_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_133, c_0_113]), c_0_113]), c_0_131]), c_0_131])).
cnf(c_0_136, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_134, c_0_135])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 137
# Proof object clause steps            : 103
# Proof object formula steps           : 34
# Proof object conjectures             : 40
# Proof object clause conjectures      : 37
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 56
# Proof object simplifying inferences  : 89
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 36
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 36
# Processed clauses                    : 1211
# ...of these trivial                  : 27
# ...subsumed                          : 668
# ...remaining for further processing  : 516
# Other redundant clauses eliminated   : 57
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 84
# Backward-rewritten                   : 327
# Generated clauses                    : 3018
# ...of the previous two non-trivial   : 2906
# Contextual simplify-reflections      : 23
# Paramodulations                      : 2953
# Factorizations                       : 9
# Equation resolutions                 : 57
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 96
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 78
# Current number of unprocessed clauses: 1448
# ...number of literals in the above   : 5928
# Current number of archived formulas  : 0
# Current number of archived clauses   : 411
# Clause-clause subsumption calls (NU) : 32922
# Rec. Clause-clause subsumption calls : 15442
# Non-unit clause-clause subsumptions  : 726
# Unit Clause-clause subsumption calls : 332
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 66
# BW rewrite match successes           : 25
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 36360

# -------------------------------------------------
# User time                : 0.049 s
# System time              : 0.000 s
# Total time               : 0.049 s
# Maximum resident set size: 28720 pages
