# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t38_wellord1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>((v2_wellord1(X3)&(r2_hidden(X1,k3_relat_1(X3))&r2_hidden(X2,k3_relat_1(X3))))=>(r1_tarski(k1_wellord1(X3,X1),k1_wellord1(X3,X2))<=>(X1=X2|r2_hidden(X1,k1_wellord1(X3,X2)))))), file('data/MPTP/proofs/0631e1f6cb5b4c229aa683330bda5621.E_input', t38_wellord1)).
fof(t37_wellord1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>((v2_wellord1(X3)&(r2_hidden(X1,k3_relat_1(X3))&r2_hidden(X2,k3_relat_1(X3))))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>r1_tarski(k1_wellord1(X3,X1),k1_wellord1(X3,X2))))), file('data/MPTP/proofs/0631e1f6cb5b4c229aa683330bda5621.E_input', t37_wellord1)).
fof(d1_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(X3=k1_wellord1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(~(X4=X2)&r2_hidden(k4_tarski(X4,X2),X1))))), file('data/MPTP/proofs/0631e1f6cb5b4c229aa683330bda5621.E_input', d1_wellord1)).
fof(l1_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v1_relat_2(X1)<=>![X2]:(r2_hidden(X2,k3_relat_1(X1))=>r2_hidden(k4_tarski(X2,X2),X1)))), file('data/MPTP/proofs/0631e1f6cb5b4c229aa683330bda5621.E_input', l1_wellord1)).
fof(d4_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v2_wellord1(X1)<=>(v1_relat_2(X1)&(v8_relat_2(X1)&(v4_relat_2(X1)&(v6_relat_2(X1)&v1_wellord1(X1))))))), file('data/MPTP/proofs/0631e1f6cb5b4c229aa683330bda5621.E_input', d4_wellord1)).
fof(c_0_5, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>((v2_wellord1(X3)&(r2_hidden(X1,k3_relat_1(X3))&r2_hidden(X2,k3_relat_1(X3))))=>(r1_tarski(k1_wellord1(X3,X1),k1_wellord1(X3,X2))<=>(X1=X2|r2_hidden(X1,k1_wellord1(X3,X2))))))), inference(assume_negation,[status(cth)],[t38_wellord1])).
fof(c_0_6, negated_conjecture, (v1_relat_1(esk3_0)&((v2_wellord1(esk3_0)&(r2_hidden(esk1_0,k3_relat_1(esk3_0))&r2_hidden(esk2_0,k3_relat_1(esk3_0))))&(((esk1_0!=esk2_0|~r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0)))&(~r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))|~r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))))&(r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))|(esk1_0=esk2_0|r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
fof(c_0_7, plain, ![X11, X12, X13]:((~r2_hidden(k4_tarski(X11,X12),X13)|r1_tarski(k1_wellord1(X13,X11),k1_wellord1(X13,X12))|(~v2_wellord1(X13)|(~r2_hidden(X11,k3_relat_1(X13))|~r2_hidden(X12,k3_relat_1(X13))))|~v1_relat_1(X13))&(~r1_tarski(k1_wellord1(X13,X11),k1_wellord1(X13,X12))|r2_hidden(k4_tarski(X11,X12),X13)|(~v2_wellord1(X13)|(~r2_hidden(X11,k3_relat_1(X13))|~r2_hidden(X12,k3_relat_1(X13))))|~v1_relat_1(X13))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t37_wellord1])])])).
fof(c_0_8, plain, ![X14, X15, X16, X17, X18, X19, X20]:((((X17!=X15|~r2_hidden(X17,X16)|X16!=k1_wellord1(X14,X15)|~v1_relat_1(X14))&(r2_hidden(k4_tarski(X17,X15),X14)|~r2_hidden(X17,X16)|X16!=k1_wellord1(X14,X15)|~v1_relat_1(X14)))&(X18=X15|~r2_hidden(k4_tarski(X18,X15),X14)|r2_hidden(X18,X16)|X16!=k1_wellord1(X14,X15)|~v1_relat_1(X14)))&((~r2_hidden(esk5_3(X14,X19,X20),X20)|(esk5_3(X14,X19,X20)=X19|~r2_hidden(k4_tarski(esk5_3(X14,X19,X20),X19),X14))|X20=k1_wellord1(X14,X19)|~v1_relat_1(X14))&((esk5_3(X14,X19,X20)!=X19|r2_hidden(esk5_3(X14,X19,X20),X20)|X20=k1_wellord1(X14,X19)|~v1_relat_1(X14))&(r2_hidden(k4_tarski(esk5_3(X14,X19,X20),X19),X14)|r2_hidden(esk5_3(X14,X19,X20),X20)|X20=k1_wellord1(X14,X19)|~v1_relat_1(X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_wellord1])])])])])])])).
cnf(c_0_9, negated_conjecture, (~r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))|~r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (r1_tarski(k1_wellord1(X3,X1),k1_wellord1(X3,X2))|~r2_hidden(k4_tarski(X1,X2),X3)|~v2_wellord1(X3)|~r2_hidden(X1,k3_relat_1(X3))|~r2_hidden(X2,k3_relat_1(X3))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, negated_conjecture, (r2_hidden(esk2_0,k3_relat_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12, negated_conjecture, (r2_hidden(esk1_0,k3_relat_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13, negated_conjecture, (v2_wellord1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,X4)|X4!=k1_wellord1(X3,X2)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16, plain, (r2_hidden(k4_tarski(X2,X3),X1)|~r1_tarski(k1_wellord1(X1,X2),k1_wellord1(X1,X3))|~v2_wellord1(X1)|~r2_hidden(X2,k3_relat_1(X1))|~r2_hidden(X3,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_17, negated_conjecture, (r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))|esk1_0=esk2_0|r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_18, negated_conjecture, (~r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))|~r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11]), c_0_12]), c_0_13]), c_0_14])])).
cnf(c_0_19, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,k1_wellord1(X3,X2))|~v1_relat_1(X3)), inference(er,[status(thm)],[c_0_15])).
cnf(c_0_20, plain, (X1=X2|r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_wellord1(X3,X2)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_21, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_11]), c_0_12]), c_0_13]), c_0_14])])).
cnf(c_0_22, negated_conjecture, (~r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_14])])).
cnf(c_0_23, negated_conjecture, (esk1_0!=esk2_0|~r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_24, plain, (X1=X2|r2_hidden(X1,k1_wellord1(X3,X2))|~r2_hidden(k4_tarski(X1,X2),X3)|~v1_relat_1(X3)), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_25, negated_conjecture, (esk2_0=esk1_0|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(sr,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_26, negated_conjecture, (esk2_0!=esk1_0|~r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_10]), c_0_11]), c_0_12]), c_0_13]), c_0_14])])).
cnf(c_0_27, negated_conjecture, (esk2_0=esk1_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_14])]), c_0_22])).
fof(c_0_28, plain, ![X8, X9]:((~v1_relat_2(X8)|(~r2_hidden(X9,k3_relat_1(X8))|r2_hidden(k4_tarski(X9,X9),X8))|~v1_relat_1(X8))&((r2_hidden(esk4_1(X8),k3_relat_1(X8))|v1_relat_2(X8)|~v1_relat_1(X8))&(~r2_hidden(k4_tarski(esk4_1(X8),esk4_1(X8)),X8)|v1_relat_2(X8)|~v1_relat_1(X8)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_wellord1])])])])])).
cnf(c_0_29, negated_conjecture, (~r2_hidden(k4_tarski(esk1_0,esk1_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_27]), c_0_27])])).
cnf(c_0_30, plain, (r2_hidden(k4_tarski(X2,X2),X1)|~v1_relat_2(X1)|~r2_hidden(X2,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_31, plain, ![X22]:(((v1_relat_2(X22)|~v2_wellord1(X22)|~v1_relat_1(X22))&((v8_relat_2(X22)|~v2_wellord1(X22)|~v1_relat_1(X22))&((v4_relat_2(X22)|~v2_wellord1(X22)|~v1_relat_1(X22))&((v6_relat_2(X22)|~v2_wellord1(X22)|~v1_relat_1(X22))&(v1_wellord1(X22)|~v2_wellord1(X22)|~v1_relat_1(X22))))))&(~v1_relat_2(X22)|(~v8_relat_2(X22)|(~v4_relat_2(X22)|(~v6_relat_2(X22)|~v1_wellord1(X22))))|v2_wellord1(X22)|~v1_relat_1(X22))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_wellord1])])])).
cnf(c_0_32, negated_conjecture, (~v1_relat_2(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_12]), c_0_14])])).
cnf(c_0_33, plain, (v1_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_34, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_13]), c_0_14])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 35
# Proof object clause steps            : 24
# Proof object formula steps           : 11
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 32
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 24
# Processed clauses                    : 43
# ...of these trivial                  : 0
# ...subsumed                          : 5
# ...remaining for further processing  : 38
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 8
# Generated clauses                    : 33
# ...of the previous two non-trivial   : 27
# Contextual simplify-reflections      : 0
# Paramodulations                      : 28
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 25
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 17
# Current number of unprocessed clauses: 8
# ...number of literals in the above   : 21
# Current number of archived formulas  : 0
# Current number of archived clauses   : 10
# Clause-clause subsumption calls (NU) : 195
# Rec. Clause-clause subsumption calls : 46
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 6
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2073

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.003 s
# Total time               : 0.009 s
# Maximum resident set size: 28656 pages
