# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t134_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4)))), file('data/MPTP/proofs/063d9c6564f24b64a17e7252cf53f32a.E_input', t134_zfmisc_1)).
fof(t120_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&k2_zfmisc_1(X3,k2_xboole_0(X1,X2))=k2_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))), file('data/MPTP/proofs/063d9c6564f24b64a17e7252cf53f32a.E_input', t120_zfmisc_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/063d9c6564f24b64a17e7252cf53f32a.E_input', commutativity_k2_xboole_0)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/063d9c6564f24b64a17e7252cf53f32a.E_input', t7_xboole_1)).
fof(t23_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3)), file('data/MPTP/proofs/063d9c6564f24b64a17e7252cf53f32a.E_input', t23_xboole_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/063d9c6564f24b64a17e7252cf53f32a.E_input', idempotence_k3_xboole_0)).
fof(t21_xboole_1, axiom, ![X1, X2]:k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/063d9c6564f24b64a17e7252cf53f32a.E_input', t21_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/063d9c6564f24b64a17e7252cf53f32a.E_input', commutativity_k3_xboole_0)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/063d9c6564f24b64a17e7252cf53f32a.E_input', d10_xboole_0)).
fof(t117_zfmisc_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&((r1_tarski(k2_zfmisc_1(X2,X1),k2_zfmisc_1(X3,X1))|r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))&~(r1_tarski(X2,X3))))), file('data/MPTP/proofs/063d9c6564f24b64a17e7252cf53f32a.E_input', t117_zfmisc_1)).
fof(t123_zfmisc_1, axiom, ![X1, X2, X3, X4]:k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)), file('data/MPTP/proofs/063d9c6564f24b64a17e7252cf53f32a.E_input', t123_zfmisc_1)).
fof(c_0_11, negated_conjecture, ~(![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4))))), inference(assume_negation,[status(cth)],[t134_zfmisc_1])).
fof(c_0_12, plain, ![X26, X27, X28]:(k2_zfmisc_1(k2_xboole_0(X26,X27),X28)=k2_xboole_0(k2_zfmisc_1(X26,X28),k2_zfmisc_1(X27,X28))&k2_zfmisc_1(X28,k2_xboole_0(X26,X27))=k2_xboole_0(k2_zfmisc_1(X28,X26),k2_zfmisc_1(X28,X27))), inference(variable_rename,[status(thm)],[t120_zfmisc_1])).
fof(c_0_13, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k2_zfmisc_1(esk3_0,esk4_0)&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk1_0!=esk3_0|esk2_0!=esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
fof(c_0_14, plain, ![X31, X32]:k2_xboole_0(X31,X32)=k2_xboole_0(X32,X31), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_15, plain, ![X17, X18]:r1_tarski(X17,k2_xboole_0(X17,X18)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_16, plain, ![X13, X14, X15]:k3_xboole_0(X13,k2_xboole_0(X14,X15))=k2_xboole_0(k3_xboole_0(X13,X14),k3_xboole_0(X13,X15)), inference(variable_rename,[status(thm)],[t23_xboole_1])).
fof(c_0_17, plain, ![X16]:k3_xboole_0(X16,X16)=X16, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
fof(c_0_18, plain, ![X22, X23]:k3_xboole_0(X22,k2_xboole_0(X22,X23))=X22, inference(variable_rename,[status(thm)],[t21_xboole_1])).
cnf(c_0_19, plain, (k2_zfmisc_1(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k2_zfmisc_1(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (k3_xboole_0(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_26, plain, ![X24, X25]:k3_xboole_0(X24,X25)=k3_xboole_0(X25,X24), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_27, negated_conjecture, (k2_xboole_0(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,X1))=k2_zfmisc_1(esk3_0,k2_xboole_0(X1,esk4_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])).
cnf(c_0_28, plain, (k2_zfmisc_1(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_29, plain, ![X29, X30]:(((r1_tarski(X29,X30)|X29!=X30)&(r1_tarski(X30,X29)|X29!=X30))&(~r1_tarski(X29,X30)|~r1_tarski(X30,X29)|X29=X30)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_30, plain, (r1_tarski(k3_xboole_0(X1,X2),k3_xboole_0(X1,k2_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_31, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_22, c_0_21])).
cnf(c_0_32, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
cnf(c_0_33, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_34, plain, ![X19, X20, X21]:((~r1_tarski(k2_zfmisc_1(X20,X19),k2_zfmisc_1(X21,X19))|r1_tarski(X20,X21)|X19=k1_xboole_0)&(~r1_tarski(k2_zfmisc_1(X19,X20),k2_zfmisc_1(X19,X21))|r1_tarski(X20,X21)|X19=k1_xboole_0)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t117_zfmisc_1])])])])).
fof(c_0_35, plain, ![X9, X10, X11, X12]:k2_zfmisc_1(k3_xboole_0(X9,X10),k3_xboole_0(X11,X12))=k3_xboole_0(k2_zfmisc_1(X9,X11),k2_zfmisc_1(X10,X12)), inference(variable_rename,[status(thm)],[t123_zfmisc_1])).
cnf(c_0_36, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,k2_xboole_0(X1,esk4_0)))=k2_zfmisc_1(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_25, c_0_27])).
cnf(c_0_37, negated_conjecture, (k2_zfmisc_1(k2_xboole_0(esk1_0,esk3_0),esk2_0)=k2_zfmisc_1(esk3_0,k2_xboole_0(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_28, c_0_27])).
cnf(c_0_38, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, plain, (r1_tarski(X1,k3_xboole_0(X1,k2_xboole_0(k2_xboole_0(X1,X2),X3)))), inference(spm,[status(thm)],[c_0_30, c_0_25])).
cnf(c_0_40, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_41, plain, (k2_xboole_0(X1,k3_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_42, plain, (k3_xboole_0(X1,k2_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_25, c_0_21])).
cnf(c_0_43, plain, (r1_tarski(X1,X3)|X2=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_44, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_45, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk1_0,esk2_0),k2_xboole_0(X1,k2_zfmisc_1(esk3_0,k2_xboole_0(X2,esk4_0))))=k2_zfmisc_1(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_36]), c_0_21]), c_0_32])).
cnf(c_0_46, negated_conjecture, (k2_xboole_0(k2_zfmisc_1(X1,esk2_0),k2_zfmisc_1(esk3_0,k2_xboole_0(esk2_0,esk4_0)))=k2_zfmisc_1(k2_xboole_0(X1,k2_xboole_0(esk1_0,esk3_0)),esk2_0)), inference(spm,[status(thm)],[c_0_28, c_0_37])).
cnf(c_0_47, plain, (k3_xboole_0(X1,k2_xboole_0(k2_xboole_0(X1,X2),X3))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])])).
cnf(c_0_48, plain, (k2_xboole_0(X1,k2_xboole_0(X2,X1))=k2_xboole_0(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_21])).
cnf(c_0_49, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|r1_tarski(k3_xboole_0(X3,X4),X5)|~r1_tarski(k3_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X4,X2)),k2_zfmisc_1(X5,k3_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_50, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(k2_xboole_0(X1,k2_xboole_0(esk1_0,esk3_0)),esk2_0))=k2_zfmisc_1(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_51, plain, (k3_xboole_0(X1,k2_xboole_0(X2,k2_xboole_0(X1,X3)))=X1), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_52, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_53, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,k2_xboole_0(X2,X4)))=k2_zfmisc_1(k3_xboole_0(X1,X3),X2)), inference(spm,[status(thm)],[c_0_44, c_0_25])).
cnf(c_0_54, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_55, negated_conjecture, (r1_tarski(esk1_0,X1)|~r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(X1,esk2_0))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_24]), c_0_51]), c_0_24]), c_0_52])).
cnf(c_0_56, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk1_0,esk3_0),esk2_0)=k2_zfmisc_1(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_36, c_0_53])).
cnf(c_0_57, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_54])).
cnf(c_0_58, negated_conjecture, (r1_tarski(esk1_0,k3_xboole_0(esk1_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57])])).
cnf(c_0_59, negated_conjecture, (k2_xboole_0(k2_zfmisc_1(X1,esk4_0),k2_zfmisc_1(esk1_0,esk2_0))=k2_zfmisc_1(k2_xboole_0(X1,esk3_0),esk4_0)), inference(spm,[status(thm)],[c_0_28, c_0_20])).
cnf(c_0_60, plain, (r1_tarski(X2,X3)|X1=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_61, negated_conjecture, (k3_xboole_0(esk1_0,esk3_0)=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_58]), c_0_40])])).
cnf(c_0_62, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(k2_xboole_0(X1,X3),X2))=k2_zfmisc_1(X1,X2)), inference(spm,[status(thm)],[c_0_25, c_0_28])).
cnf(c_0_63, negated_conjecture, (k2_zfmisc_1(k2_xboole_0(esk1_0,esk3_0),esk4_0)=k2_zfmisc_1(esk1_0,k2_xboole_0(esk2_0,esk4_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_59]), c_0_21])).
cnf(c_0_64, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|r1_tarski(X3,k3_xboole_0(X4,X5))|~r1_tarski(k2_zfmisc_1(k3_xboole_0(X1,X2),X3),k3_xboole_0(k2_zfmisc_1(X1,X4),k2_zfmisc_1(X2,X5)))), inference(spm,[status(thm)],[c_0_60, c_0_44])).
cnf(c_0_65, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_66, negated_conjecture, (k2_xboole_0(esk1_0,esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_61]), c_0_21])).
cnf(c_0_67, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk1_0,esk4_0),k2_zfmisc_1(esk1_0,k2_xboole_0(esk2_0,esk4_0)))=k2_zfmisc_1(esk1_0,esk4_0)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_68, negated_conjecture, (r1_tarski(X1,esk2_0)|~r1_tarski(k2_zfmisc_1(esk1_0,X1),k2_zfmisc_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_50]), c_0_51]), c_0_24]), c_0_51]), c_0_65])).
cnf(c_0_69, negated_conjecture, (k2_zfmisc_1(esk1_0,k2_xboole_0(esk2_0,esk4_0))=k2_zfmisc_1(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_66]), c_0_20])).
cnf(c_0_70, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk1_0,esk4_0),k2_xboole_0(k2_zfmisc_1(esk1_0,k2_xboole_0(esk2_0,esk4_0)),X1))=k2_zfmisc_1(esk1_0,esk4_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_67]), c_0_32])).
cnf(c_0_71, negated_conjecture, (r1_tarski(k2_xboole_0(esk2_0,esk4_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_57])])).
cnf(c_0_72, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk1_0,esk4_0),k2_zfmisc_1(k2_xboole_0(esk1_0,X1),k2_xboole_0(esk2_0,esk4_0)))=k2_zfmisc_1(esk1_0,esk4_0)), inference(spm,[status(thm)],[c_0_70, c_0_28])).
cnf(c_0_73, negated_conjecture, (k2_xboole_0(esk2_0,esk4_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_71]), c_0_22])])).
cnf(c_0_74, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk1_0,esk4_0),k2_zfmisc_1(k2_xboole_0(esk1_0,X1),esk2_0))=k2_zfmisc_1(esk1_0,esk4_0)), inference(rw,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_75, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|r1_tarski(k3_xboole_0(X3,X4),X5)|~r1_tarski(k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)),k2_zfmisc_1(k3_xboole_0(X1,X2),X5))), inference(spm,[status(thm)],[c_0_60, c_0_44])).
cnf(c_0_76, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk1_0,esk4_0),k2_zfmisc_1(esk3_0,esk2_0))=k2_zfmisc_1(esk1_0,esk4_0)), inference(spm,[status(thm)],[c_0_74, c_0_66])).
cnf(c_0_77, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X4),k2_zfmisc_1(X3,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_33]), c_0_44])).
cnf(c_0_78, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(k2_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_38, c_0_31])).
cnf(c_0_79, negated_conjecture, (r1_tarski(esk2_0,X1)|~r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk1_0,X1))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_50]), c_0_51]), c_0_24]), c_0_51]), c_0_65])).
cnf(c_0_80, negated_conjecture, (k2_zfmisc_1(esk1_0,esk4_0)=k2_zfmisc_1(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_77]), c_0_20]), c_0_24])).
cnf(c_0_81, negated_conjecture, (esk4_0=esk2_0|~r1_tarski(esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_78, c_0_73])).
cnf(c_0_82, negated_conjecture, (r1_tarski(esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_57])])).
cnf(c_0_83, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_40, c_0_33])).
cnf(c_0_84, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|r1_tarski(X3,k3_xboole_0(X4,X5))|~r1_tarski(k2_zfmisc_1(X3,k3_xboole_0(X1,X2)),k3_xboole_0(k2_zfmisc_1(X4,X1),k2_zfmisc_1(X5,X2)))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_85, negated_conjecture, (esk4_0=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_82])])).
cnf(c_0_86, plain, (k3_xboole_0(X1,X2)=X2|~r1_tarski(X2,k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_38, c_0_83])).
cnf(c_0_87, negated_conjecture, (r1_tarski(X1,esk1_0)|~r1_tarski(k2_zfmisc_1(X1,esk2_0),k2_zfmisc_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_50]), c_0_24]), c_0_51]), c_0_24]), c_0_52])).
cnf(c_0_88, negated_conjecture, (k2_zfmisc_1(esk3_0,esk2_0)=k2_zfmisc_1(esk1_0,esk2_0)), inference(rw,[status(thm)],[c_0_20, c_0_85])).
cnf(c_0_89, negated_conjecture, (esk1_0!=esk3_0|esk2_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_90, negated_conjecture, (esk3_0=esk1_0|~r1_tarski(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_86, c_0_61])).
cnf(c_0_91, negated_conjecture, (r1_tarski(esk3_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_57])])).
cnf(c_0_92, negated_conjecture, (esk3_0!=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_85])])).
cnf(c_0_93, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_91])]), c_0_92]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 94
# Proof object clause steps            : 71
# Proof object formula steps           : 23
# Proof object conjectures             : 40
# Proof object clause conjectures      : 37
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 46
# Proof object simplifying inferences  : 54
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 18
# Processed clauses                    : 729
# ...of these trivial                  : 261
# ...subsumed                          : 205
# ...remaining for further processing  : 263
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 140
# Generated clauses                    : 5316
# ...of the previous two non-trivial   : 3668
# Contextual simplify-reflections      : 0
# Paramodulations                      : 5314
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 104
#    Positive orientable unit clauses  : 75
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 23
# Current number of unprocessed clauses: 2936
# ...number of literals in the above   : 4096
# Current number of archived formulas  : 0
# Current number of archived clauses   : 157
# Clause-clause subsumption calls (NU) : 761
# Rec. Clause-clause subsumption calls : 761
# Non-unit clause-clause subsumptions  : 196
# Unit Clause-clause subsumption calls : 23
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 206
# BW rewrite match successes           : 93
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 64015

# -------------------------------------------------
# User time                : 0.033 s
# System time              : 0.003 s
# Total time               : 0.037 s
# Maximum resident set size: 28436 pages
