# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t68_mcart_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&?[X5, X6, X7]:(X4=k3_mcart_1(X5,X6,X7)&~((k5_mcart_1(X1,X2,X3,X4)=X5&(k6_mcart_1(X1,X2,X3,X4)=X6&k7_mcart_1(X1,X2,X3,X4)=X7))))))))), file('data/MPTP/proofs/06460cb64cc44fb0be3ce49b5da93040.E_input', t68_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/06460cb64cc44fb0be3ce49b5da93040.E_input', d3_zfmisc_1)).
fof(t50_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))&(k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))&k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4))))))))), file('data/MPTP/proofs/06460cb64cc44fb0be3ce49b5da93040.E_input', t50_mcart_1)).
fof(l44_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&?[X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))&![X5]:(m1_subset_1(X5,X1)=>![X6]:(m1_subset_1(X6,X2)=>![X7]:(m1_subset_1(X7,X3)=>~(X4=k3_mcart_1(X5,X6,X7)))))))))), file('data/MPTP/proofs/06460cb64cc44fb0be3ce49b5da93040.E_input', l44_mcart_1)).
fof(t71_mcart_1, conjecture, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X8))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k7_mcart_1(X1,X2,X3,X5)))))), file('data/MPTP/proofs/06460cb64cc44fb0be3ce49b5da93040.E_input', t71_mcart_1)).
fof(c_0_5, plain, ![X9, X10, X11, X12, X13, X14, X15]:((k5_mcart_1(X9,X10,X11,X12)=X13|X12!=k3_mcart_1(X13,X14,X15)|X11=k1_xboole_0|X10=k1_xboole_0|X9=k1_xboole_0|~m1_subset_1(X12,k3_zfmisc_1(X9,X10,X11)))&((k6_mcart_1(X9,X10,X11,X12)=X14|X12!=k3_mcart_1(X13,X14,X15)|X11=k1_xboole_0|X10=k1_xboole_0|X9=k1_xboole_0|~m1_subset_1(X12,k3_zfmisc_1(X9,X10,X11)))&(k7_mcart_1(X9,X10,X11,X12)=X15|X12!=k3_mcart_1(X13,X14,X15)|X11=k1_xboole_0|X10=k1_xboole_0|X9=k1_xboole_0|~m1_subset_1(X12,k3_zfmisc_1(X9,X10,X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t68_mcart_1])])])])])).
fof(c_0_6, plain, ![X16, X17, X18]:k3_zfmisc_1(X16,X17,X18)=k2_zfmisc_1(k2_zfmisc_1(X16,X17),X18), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_7, plain, ![X19, X20, X21, X22]:((k5_mcart_1(X19,X20,X21,X22)=k1_mcart_1(k1_mcart_1(X22))|~m1_subset_1(X22,k3_zfmisc_1(X19,X20,X21))|X21=k1_xboole_0|X20=k1_xboole_0|X19=k1_xboole_0)&((k6_mcart_1(X19,X20,X21,X22)=k2_mcart_1(k1_mcart_1(X22))|~m1_subset_1(X22,k3_zfmisc_1(X19,X20,X21))|X21=k1_xboole_0|X20=k1_xboole_0|X19=k1_xboole_0)&(k7_mcart_1(X19,X20,X21,X22)=k2_mcart_1(X22)|~m1_subset_1(X22,k3_zfmisc_1(X19,X20,X21))|X21=k1_xboole_0|X20=k1_xboole_0|X19=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_mcart_1])])])])])).
cnf(c_0_8, plain, (k5_mcart_1(X1,X2,X3,X4)=X5|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X4!=k3_mcart_1(X5,X6,X7)|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k5_mcart_1(X1,X2,X3,X4)=X5|X4!=k3_mcart_1(X5,X6,X7)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_8, c_0_9])).
fof(c_0_12, plain, ![X23, X24, X25, X26]:((m1_subset_1(esk1_4(X23,X24,X25,X26),X23)|~m1_subset_1(X26,k3_zfmisc_1(X23,X24,X25))|X25=k1_xboole_0|X24=k1_xboole_0|X23=k1_xboole_0)&((m1_subset_1(esk2_4(X23,X24,X25,X26),X24)|~m1_subset_1(X26,k3_zfmisc_1(X23,X24,X25))|X25=k1_xboole_0|X24=k1_xboole_0|X23=k1_xboole_0)&((m1_subset_1(esk3_4(X23,X24,X25,X26),X25)|~m1_subset_1(X26,k3_zfmisc_1(X23,X24,X25))|X25=k1_xboole_0|X24=k1_xboole_0|X23=k1_xboole_0)&(X26=k3_mcart_1(esk1_4(X23,X24,X25,X26),esk2_4(X23,X24,X25,X26),esk3_4(X23,X24,X25,X26))|~m1_subset_1(X26,k3_zfmisc_1(X23,X24,X25))|X25=k1_xboole_0|X24=k1_xboole_0|X23=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l44_mcart_1])])])])])])])).
fof(c_0_13, negated_conjecture, ~(![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X8))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k7_mcart_1(X1,X2,X3,X5))))))), inference(assume_negation,[status(cth)],[t71_mcart_1])).
cnf(c_0_14, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_10, c_0_9])).
cnf(c_0_15, plain, (k5_mcart_1(X1,X2,X3,k3_mcart_1(X4,X5,X6))=X4|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(k3_mcart_1(X4,X5,X6),k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(er,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (X1=k3_mcart_1(esk1_4(X2,X3,X4,X1),esk2_4(X2,X3,X4,X1),esk3_4(X2,X3,X4,X1))|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|~m1_subset_1(X1,k3_zfmisc_1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_17, negated_conjecture, ![X35, X36, X37]:(m1_subset_1(esk8_0,k3_zfmisc_1(esk4_0,esk5_0,esk6_0))&((~m1_subset_1(X35,esk4_0)|(~m1_subset_1(X36,esk5_0)|(~m1_subset_1(X37,esk6_0)|(esk8_0!=k3_mcart_1(X35,X36,X37)|esk7_0=X37))))&(esk4_0!=k1_xboole_0&(esk5_0!=k1_xboole_0&(esk6_0!=k1_xboole_0&esk7_0!=k7_mcart_1(esk4_0,esk5_0,esk6_0,esk8_0)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])).
cnf(c_0_18, plain, (k1_mcart_1(k1_mcart_1(k3_mcart_1(X1,X2,X3)))=X1|X4=k1_xboole_0|X5=k1_xboole_0|X6=k1_xboole_0|~m1_subset_1(k3_mcart_1(X1,X2,X3),k2_zfmisc_1(k2_zfmisc_1(X4,X5),X6))), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_19, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k3_mcart_1(esk1_4(X2,X3,X4,X1),esk2_4(X2,X3,X4,X1),esk3_4(X2,X3,X4,X1))|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4))), inference(rw,[status(thm)],[c_0_16, c_0_9])).
cnf(c_0_20, negated_conjecture, (m1_subset_1(esk8_0,k3_zfmisc_1(esk4_0,esk5_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21, plain, (m1_subset_1(esk1_4(X1,X2,X3,X4),X1)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22, plain, (k1_mcart_1(k1_mcart_1(X1))=esk1_4(X2,X3,X4,X1)|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|X5=k1_xboole_0|X6=k1_xboole_0|X7=k1_xboole_0|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(X7,X6),X5))|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_23, negated_conjecture, (m1_subset_1(esk8_0,k2_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0),esk6_0))), inference(rw,[status(thm)],[c_0_20, c_0_9])).
cnf(c_0_24, negated_conjecture, (esk6_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, negated_conjecture, (esk5_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, negated_conjecture, (esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|m1_subset_1(esk1_4(X1,X2,X3,X4),X1)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_21, c_0_9])).
cnf(c_0_28, negated_conjecture, (k1_mcart_1(k1_mcart_1(esk8_0))=esk1_4(X1,X2,X3,esk8_0)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk8_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_25]), c_0_26])).
cnf(c_0_29, negated_conjecture, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|m1_subset_1(k1_mcart_1(k1_mcart_1(esk8_0)),X1)|~m1_subset_1(esk8_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_30, negated_conjecture, (m1_subset_1(k1_mcart_1(k1_mcart_1(esk8_0)),esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_23]), c_0_26]), c_0_25]), c_0_24])).
cnf(c_0_31, negated_conjecture, (esk7_0=X3|~m1_subset_1(X1,esk4_0)|~m1_subset_1(X2,esk5_0)|~m1_subset_1(X3,esk6_0)|esk8_0!=k3_mcart_1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_32, negated_conjecture, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|m1_subset_1(esk1_4(X1,X2,X3,esk8_0),esk4_0)|~m1_subset_1(esk8_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(spm,[status(thm)],[c_0_30, c_0_28])).
cnf(c_0_33, plain, (m1_subset_1(esk3_4(X1,X2,X3,X4),X3)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_34, negated_conjecture, (esk3_4(X1,X2,X3,esk8_0)=esk7_0|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|~m1_subset_1(esk3_4(X1,X2,X3,esk8_0),esk6_0)|~m1_subset_1(esk2_4(X1,X2,X3,esk8_0),esk5_0)|~m1_subset_1(esk8_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_19])]), c_0_32])).
cnf(c_0_35, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|m1_subset_1(esk3_4(X1,X2,X3,X4),X3)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_33, c_0_9])).
cnf(c_0_36, plain, (m1_subset_1(esk2_4(X1,X2,X3,X4),X2)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_37, plain, (k7_mcart_1(X1,X2,X3,X4)=X5|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X4!=k3_mcart_1(X6,X7,X5)|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_38, negated_conjecture, (esk3_4(X1,X2,esk6_0,esk8_0)=esk7_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk2_4(X1,X2,esk6_0,esk8_0),esk5_0)|~m1_subset_1(esk8_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),esk6_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_24])).
cnf(c_0_39, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|m1_subset_1(esk2_4(X1,X2,X3,X4),X2)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_36, c_0_9])).
cnf(c_0_40, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k7_mcart_1(X1,X2,X3,X4)=X5|X4!=k3_mcart_1(X6,X7,X5)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_37, c_0_9])).
cnf(c_0_41, negated_conjecture, (esk3_4(X1,esk5_0,esk6_0,esk8_0)=esk7_0|X1=k1_xboole_0|~m1_subset_1(esk8_0,k2_zfmisc_1(k2_zfmisc_1(X1,esk5_0),esk6_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_25]), c_0_24])).
cnf(c_0_42, plain, (k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_43, plain, (k7_mcart_1(X1,X2,X3,k3_mcart_1(X4,X5,X6))=X6|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(k3_mcart_1(X4,X5,X6),k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_44, negated_conjecture, (k3_mcart_1(esk1_4(X1,esk5_0,esk6_0,esk8_0),esk2_4(X1,esk5_0,esk6_0,esk8_0),esk7_0)=esk8_0|X1=k1_xboole_0|~m1_subset_1(esk8_0,k2_zfmisc_1(k2_zfmisc_1(X1,esk5_0),esk6_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_41]), c_0_24]), c_0_25])).
cnf(c_0_45, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_42, c_0_9])).
cnf(c_0_46, negated_conjecture, (k7_mcart_1(X1,X2,X3,esk8_0)=esk7_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk8_0,k2_zfmisc_1(k2_zfmisc_1(X4,esk5_0),esk6_0))|~m1_subset_1(esk8_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_47, negated_conjecture, (esk7_0!=k7_mcart_1(esk4_0,esk5_0,esk6_0,esk8_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_48, negated_conjecture, (k7_mcart_1(esk4_0,esk5_0,esk6_0,esk8_0)=k2_mcart_1(esk8_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_23]), c_0_24]), c_0_25]), c_0_26])).
cnf(c_0_49, negated_conjecture, (k7_mcart_1(X1,X2,X3,esk8_0)=esk7_0|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|~m1_subset_1(esk8_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_23]), c_0_26])).
cnf(c_0_50, negated_conjecture, (k2_mcart_1(esk8_0)!=esk7_0), inference(rw,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_51, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_23]), c_0_48]), c_0_50]), c_0_26]), c_0_25]), c_0_24]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 52
# Proof object clause steps            : 41
# Proof object formula steps           : 11
# Proof object conjectures             : 23
# Proof object clause conjectures      : 20
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 34
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 16
# Processed clauses                    : 55
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 55
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 2
# Generated clauses                    : 55
# ...of the previous two non-trivial   : 52
# Contextual simplify-reflections      : 1
# Paramodulations                      : 51
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 33
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 25
# Current number of unprocessed clauses: 29
# ...number of literals in the above   : 243
# Current number of archived formulas  : 0
# Current number of archived clauses   : 20
# Clause-clause subsumption calls (NU) : 142
# Rec. Clause-clause subsumption calls : 11
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3069

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28768 pages
