# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C41_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d12_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(X3=k8_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>(r2_hidden(X5,X1)&r2_hidden(k4_tarski(X4,X5),X2)))))), file('data/MPTP/proofs/0663afe47ede45e19848782ae19b7645.E_input', d12_relat_1)).
fof(dt_k8_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>v1_relat_1(k8_relat_1(X1,X2))), file('data/MPTP/proofs/0663afe47ede45e19848782ae19b7645.E_input', dt_k8_relat_1)).
fof(t128_relat_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>k8_relat_1(X1,k8_relat_1(X1,X2))=k8_relat_1(X1,X2)), file('data/MPTP/proofs/0663afe47ede45e19848782ae19b7645.E_input', t128_relat_1)).
fof(c_0_3, plain, ![X10, X11, X12, X13, X14, X15, X16]:((((r2_hidden(X14,X10)|~r2_hidden(k4_tarski(X13,X14),X12)|X12!=k8_relat_1(X10,X11)|~v1_relat_1(X12)|~v1_relat_1(X11))&(r2_hidden(k4_tarski(X13,X14),X11)|~r2_hidden(k4_tarski(X13,X14),X12)|X12!=k8_relat_1(X10,X11)|~v1_relat_1(X12)|~v1_relat_1(X11)))&(~r2_hidden(X16,X10)|~r2_hidden(k4_tarski(X15,X16),X11)|r2_hidden(k4_tarski(X15,X16),X12)|X12!=k8_relat_1(X10,X11)|~v1_relat_1(X12)|~v1_relat_1(X11)))&((~r2_hidden(k4_tarski(esk3_3(X10,X11,X12),esk4_3(X10,X11,X12)),X12)|(~r2_hidden(esk4_3(X10,X11,X12),X10)|~r2_hidden(k4_tarski(esk3_3(X10,X11,X12),esk4_3(X10,X11,X12)),X11))|X12=k8_relat_1(X10,X11)|~v1_relat_1(X12)|~v1_relat_1(X11))&((r2_hidden(esk4_3(X10,X11,X12),X10)|r2_hidden(k4_tarski(esk3_3(X10,X11,X12),esk4_3(X10,X11,X12)),X12)|X12=k8_relat_1(X10,X11)|~v1_relat_1(X12)|~v1_relat_1(X11))&(r2_hidden(k4_tarski(esk3_3(X10,X11,X12),esk4_3(X10,X11,X12)),X11)|r2_hidden(k4_tarski(esk3_3(X10,X11,X12),esk4_3(X10,X11,X12)),X12)|X12=k8_relat_1(X10,X11)|~v1_relat_1(X12)|~v1_relat_1(X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_relat_1])])])])])])).
fof(c_0_4, plain, ![X8, X9]:(~v1_relat_1(X9)|v1_relat_1(k8_relat_1(X8,X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_relat_1])])).
cnf(c_0_5, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),X4)|X4!=k8_relat_1(X2,X5)|~v1_relat_1(X4)|~v1_relat_1(X5)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_6, plain, (v1_relat_1(k8_relat_1(X2,X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, plain, (r2_hidden(k4_tarski(esk3_3(X1,X2,X3),esk4_3(X1,X2,X3)),X2)|r2_hidden(k4_tarski(esk3_3(X1,X2,X3),esk4_3(X1,X2,X3)),X3)|X3=k8_relat_1(X1,X2)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_8, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k8_relat_1(X2,X4))|~v1_relat_1(X4)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_5]), c_0_6])).
cnf(c_0_9, plain, (k8_relat_1(X1,X2)=X2|r2_hidden(k4_tarski(esk3_3(X1,X2,X2),esk4_3(X1,X2,X2)),X2)|~v1_relat_1(X2)), inference(ef,[status(thm)],[c_0_7])).
fof(c_0_10, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>k8_relat_1(X1,k8_relat_1(X1,X2))=k8_relat_1(X1,X2))), inference(assume_negation,[status(cth)],[t128_relat_1])).
cnf(c_0_11, plain, (X3=k8_relat_1(X1,X2)|~r2_hidden(k4_tarski(esk3_3(X1,X2,X3),esk4_3(X1,X2,X3)),X3)|~r2_hidden(esk4_3(X1,X2,X3),X1)|~r2_hidden(k4_tarski(esk3_3(X1,X2,X3),esk4_3(X1,X2,X3)),X2)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_12, plain, (k8_relat_1(X1,k8_relat_1(X2,X3))=k8_relat_1(X2,X3)|r2_hidden(esk4_3(X1,k8_relat_1(X2,X3),k8_relat_1(X2,X3)),X2)|~v1_relat_1(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_6])).
fof(c_0_13, negated_conjecture, (v1_relat_1(esk2_0)&k8_relat_1(esk1_0,k8_relat_1(esk1_0,esk2_0))!=k8_relat_1(esk1_0,esk2_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
cnf(c_0_14, plain, (k8_relat_1(X1,k8_relat_1(X1,X2))=k8_relat_1(X1,X2)|~r2_hidden(k4_tarski(esk3_3(X1,k8_relat_1(X1,X2),k8_relat_1(X1,X2)),esk4_3(X1,k8_relat_1(X1,X2),k8_relat_1(X1,X2))),k8_relat_1(X1,X2))|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_6])).
cnf(c_0_15, negated_conjecture, (k8_relat_1(esk1_0,k8_relat_1(esk1_0,esk2_0))!=k8_relat_1(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_16, plain, (k8_relat_1(X1,k8_relat_1(X1,X2))=k8_relat_1(X1,X2)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_9]), c_0_6])).
cnf(c_0_17, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 19
# Proof object clause steps            : 12
# Proof object formula steps           : 7
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 6
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 6
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 9
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 9
# Processed clauses                    : 16
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 16
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 29
# ...of the previous two non-trivial   : 20
# Contextual simplify-reflections      : 6
# Paramodulations                      : 24
# Factorizations                       : 2
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 15
#    Positive orientable unit clauses  : 1
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 13
# Current number of unprocessed clauses: 13
# ...number of literals in the above   : 69
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1
# Clause-clause subsumption calls (NU) : 61
# Rec. Clause-clause subsumption calls : 22
# Non-unit clause-clause subsumptions  : 7
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1350

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.000 s
# Total time               : 0.009 s
# Maximum resident set size: 28716 pages
