# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d9_xboole_0, axiom, ![X1, X2]:(r3_xboole_0(X1,X2)<=>(r1_tarski(X1,X2)|r1_tarski(X2,X1))), file('data/MPTP/proofs/069ac088b5b84ecb8a9abb42ce911eff.E_input', d9_xboole_0)).
fof(d8_xboole_0, axiom, ![X1, X2]:(r2_xboole_0(X1,X2)<=>(r1_tarski(X1,X2)&~(X1=X2))), file('data/MPTP/proofs/069ac088b5b84ecb8a9abb42ce911eff.E_input', d8_xboole_0)).
fof(t104_xboole_1, conjecture, ![X1, X2]:(~((~(r2_xboole_0(X1,X2))&(~(X1=X2)&~(r2_xboole_0(X2,X1)))))<=>r3_xboole_0(X1,X2)), file('data/MPTP/proofs/069ac088b5b84ecb8a9abb42ce911eff.E_input', t104_xboole_1)).
fof(reflexivity_r3_xboole_0, axiom, ![X1, X2]:r3_xboole_0(X1,X1), file('data/MPTP/proofs/069ac088b5b84ecb8a9abb42ce911eff.E_input', reflexivity_r3_xboole_0)).
fof(c_0_4, plain, ![X3, X4]:((~r3_xboole_0(X3,X4)|(r1_tarski(X3,X4)|r1_tarski(X4,X3)))&((~r1_tarski(X3,X4)|r3_xboole_0(X3,X4))&(~r1_tarski(X4,X3)|r3_xboole_0(X3,X4)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_xboole_0])])])).
fof(c_0_5, plain, ![X5, X6]:(((r1_tarski(X5,X6)|~r2_xboole_0(X5,X6))&(X5!=X6|~r2_xboole_0(X5,X6)))&(~r1_tarski(X5,X6)|X5=X6|r2_xboole_0(X5,X6))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_xboole_0])])])])).
fof(c_0_6, negated_conjecture, ~(![X1, X2]:(~((~(r2_xboole_0(X1,X2))&(~(X1=X2)&~(r2_xboole_0(X2,X1)))))<=>r3_xboole_0(X1,X2))), inference(assume_negation,[status(cth)],[t104_xboole_1])).
cnf(c_0_7, plain, (r3_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, plain, (r1_tarski(X1,X2)|~r2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_9, negated_conjecture, (((~r2_xboole_0(esk1_0,esk2_0)|~r3_xboole_0(esk1_0,esk2_0))&((esk1_0!=esk2_0|~r3_xboole_0(esk1_0,esk2_0))&(~r2_xboole_0(esk2_0,esk1_0)|~r3_xboole_0(esk1_0,esk2_0))))&(r2_xboole_0(esk1_0,esk2_0)|(esk1_0=esk2_0|r2_xboole_0(esk2_0,esk1_0))|r3_xboole_0(esk1_0,esk2_0))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_6])])])])])).
cnf(c_0_10, plain, (r3_xboole_0(X1,X2)|~r2_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_11, negated_conjecture, (r2_xboole_0(esk1_0,esk2_0)|esk1_0=esk2_0|r2_xboole_0(esk2_0,esk1_0)|r3_xboole_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12, negated_conjecture, (esk1_0=esk2_0|r2_xboole_0(esk2_0,esk1_0)|r3_xboole_0(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_13, plain, (r1_tarski(X1,X2)|r1_tarski(X2,X1)|~r3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_14, negated_conjecture, (esk1_0=esk2_0|r3_xboole_0(esk1_0,esk2_0)|r3_xboole_0(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_10, c_0_12])).
cnf(c_0_15, plain, (X1=X2|r2_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_16, negated_conjecture, (esk1_0=esk2_0|r1_tarski(esk1_0,esk2_0)|r1_tarski(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_13])).
cnf(c_0_17, plain, (r3_xboole_0(X2,X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_18, negated_conjecture, (esk1_0=esk2_0|r2_xboole_0(esk1_0,esk2_0)|r1_tarski(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_19, plain, (r3_xboole_0(X1,X2)|~r2_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_17, c_0_8])).
cnf(c_0_20, negated_conjecture, (~r2_xboole_0(esk1_0,esk2_0)|~r3_xboole_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_21, negated_conjecture, (esk1_0=esk2_0|r2_xboole_0(esk1_0,esk2_0)|r2_xboole_0(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_15, c_0_18])).
cnf(c_0_22, negated_conjecture, (esk1_0=esk2_0|r3_xboole_0(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_19, c_0_12])).
cnf(c_0_23, negated_conjecture, (~r2_xboole_0(esk2_0,esk1_0)|~r3_xboole_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_24, negated_conjecture, (esk1_0=esk2_0|r2_xboole_0(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])).
fof(c_0_25, plain, ![X9]:r3_xboole_0(X9,X9), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r3_xboole_0])])).
cnf(c_0_26, negated_conjecture, (esk1_0!=esk2_0|~r3_xboole_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_27, negated_conjecture, (esk1_0=esk2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_22])).
cnf(c_0_28, plain, (r3_xboole_0(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_29, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_27]), c_0_27]), c_0_28])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 30
# Proof object clause steps            : 21
# Proof object formula steps           : 9
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 7
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 11
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 11
# Processed clauses                    : 37
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 35
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 8
# Generated clauses                    : 29
# ...of the previous two non-trivial   : 24
# Contextual simplify-reflections      : 4
# Paramodulations                      : 28
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 11
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 7
# Current number of unprocessed clauses: 7
# ...number of literals in the above   : 19
# Current number of archived formulas  : 0
# Current number of archived clauses   : 23
# Clause-clause subsumption calls (NU) : 33
# Rec. Clause-clause subsumption calls : 30
# Non-unit clause-clause subsumptions  : 10
# Unit Clause-clause subsumption calls : 9
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 779

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.000 s
# Total time               : 0.009 s
# Maximum resident set size: 28764 pages
