# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C41_F1_PI_AE_Q4_CS_SP_PS_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t2_yellow_1, conjecture, ![X1, X2]:(m1_subset_1(X2,u1_struct_0(k3_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k3_yellow_1(X1)))=>(r3_orders_2(k3_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3)))), file('data/MPTP/proofs/06a9bd10a54f4172a6a07cbc8983b528.E_input', t2_yellow_1)).
fof(d2_yellow_1, axiom, ![X1]:k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1)), file('data/MPTP/proofs/06a9bd10a54f4172a6a07cbc8983b528.E_input', d2_yellow_1)).
fof(d4_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(k3_lattice3(X1)))=>k5_lattice3(X1,X2)=X2)), file('data/MPTP/proofs/06a9bd10a54f4172a6a07cbc8983b528.E_input', d4_lattice3)).
fof(dt_k1_lattice3, axiom, ![X1]:(v3_lattices(k1_lattice3(X1))&l3_lattices(k1_lattice3(X1))), file('data/MPTP/proofs/06a9bd10a54f4172a6a07cbc8983b528.E_input', dt_k1_lattice3)).
fof(fc2_lattice3, axiom, ![X1]:(v3_lattices(k1_lattice3(X1))&v10_lattices(k1_lattice3(X1))), file('data/MPTP/proofs/06a9bd10a54f4172a6a07cbc8983b528.E_input', fc2_lattice3)).
fof(fc1_lattice3, axiom, ![X1]:(~(v2_struct_0(k1_lattice3(X1)))&v3_lattices(k1_lattice3(X1))), file('data/MPTP/proofs/06a9bd10a54f4172a6a07cbc8983b528.E_input', fc1_lattice3)).
fof(dt_k5_lattice3, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))&m1_subset_1(X2,u1_struct_0(k3_lattice3(X1))))=>m1_subset_1(k5_lattice3(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/06a9bd10a54f4172a6a07cbc8983b528.E_input', dt_k5_lattice3)).
fof(redefinition_r3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_lattices(X1,X2,X3)<=>r1_lattices(X1,X2,X3))), file('data/MPTP/proofs/06a9bd10a54f4172a6a07cbc8983b528.E_input', redefinition_r3_lattices)).
fof(d3_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>k4_lattice3(X1,X2)=X2)), file('data/MPTP/proofs/06a9bd10a54f4172a6a07cbc8983b528.E_input', d3_lattice3)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/06a9bd10a54f4172a6a07cbc8983b528.E_input', cc1_lattices)).
fof(t7_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_lattices(X1,X2,X3)<=>r3_orders_2(k3_lattice3(X1),k4_lattice3(X1,X2),k4_lattice3(X1,X3)))))), file('data/MPTP/proofs/06a9bd10a54f4172a6a07cbc8983b528.E_input', t7_lattice3)).
fof(t2_lattice3, axiom, ![X1, X2]:(m1_subset_1(X2,u1_struct_0(k1_lattice3(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k1_lattice3(X1)))=>(r1_lattices(k1_lattice3(X1),X2,X3)<=>r1_tarski(X2,X3)))), file('data/MPTP/proofs/06a9bd10a54f4172a6a07cbc8983b528.E_input', t2_lattice3)).
fof(c_0_12, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,u1_struct_0(k3_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k3_yellow_1(X1)))=>(r3_orders_2(k3_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3))))), inference(assume_negation,[status(cth)],[t2_yellow_1])).
fof(c_0_13, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(k3_yellow_1(esk1_0)))&((~r3_orders_2(k3_yellow_1(esk1_0),esk2_0,esk3_0)|~r1_tarski(esk2_0,esk3_0))&(r3_orders_2(k3_yellow_1(esk1_0),esk2_0,esk3_0)|r1_tarski(esk2_0,esk3_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
fof(c_0_14, plain, ![X25]:k3_yellow_1(X25)=k3_lattice3(k1_lattice3(X25)), inference(variable_rename,[status(thm)],[d2_yellow_1])).
fof(c_0_15, plain, ![X7, X8]:(v2_struct_0(X7)|(~v10_lattices(X7)|~l3_lattices(X7))|(~m1_subset_1(X8,u1_struct_0(k3_lattice3(X7)))|k5_lattice3(X7,X8)=X8)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_lattice3])])])])).
cnf(c_0_16, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k3_yellow_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, plain, (k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, plain, ![X17]:(v3_lattices(k1_lattice3(X17))&l3_lattices(k1_lattice3(X17))), inference(variable_rename,[status(thm)],[dt_k1_lattice3])).
fof(c_0_19, plain, ![X26]:(v3_lattices(k1_lattice3(X26))&v10_lattices(k1_lattice3(X26))), inference(variable_rename,[status(thm)],[fc2_lattice3])).
fof(c_0_20, plain, ![X21]:(~v2_struct_0(k1_lattice3(X21))&v3_lattices(k1_lattice3(X21))), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_lattice3])])).
fof(c_0_21, plain, ![X9, X10]:(v2_struct_0(X9)|(~v10_lattices(X9)|~l3_lattices(X9))|~m1_subset_1(X10,u1_struct_0(k3_lattice3(X9)))|m1_subset_1(k5_lattice3(X9,X10),u1_struct_0(X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_lattice3])])])).
cnf(c_0_22, plain, (v2_struct_0(X1)|k5_lattice3(X1,X2)=X2|~v10_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(k3_lattice3(X1)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k3_lattice3(k1_lattice3(esk1_0))))), inference(rw,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_24, plain, (l3_lattices(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (v10_lattices(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (~v2_struct_0(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_28, plain, ![X18, X19, X20]:((~r3_lattices(X18,X19,X20)|r1_lattices(X18,X19,X20)|(v2_struct_0(X18)|(~v6_lattices(X18)|(~v8_lattices(X18)|(~v9_lattices(X18)|~l3_lattices(X18))))|(~m1_subset_1(X19,u1_struct_0(X18))|~m1_subset_1(X20,u1_struct_0(X18)))))&(~r1_lattices(X18,X19,X20)|r3_lattices(X18,X19,X20)|(v2_struct_0(X18)|(~v6_lattices(X18)|(~v8_lattices(X18)|(~v9_lattices(X18)|~l3_lattices(X18))))|(~m1_subset_1(X19,u1_struct_0(X18))|~m1_subset_1(X20,u1_struct_0(X18)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_lattices])])])])).
cnf(c_0_29, plain, (v2_struct_0(X1)|m1_subset_1(k5_lattice3(X1,X2),u1_struct_0(X1))|~v10_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(k3_lattice3(X1)))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, negated_conjecture, (k5_lattice3(k1_lattice3(esk1_0),esk3_0)=esk3_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_25])]), c_0_26])).
fof(c_0_31, plain, ![X23, X24]:(v2_struct_0(X23)|(~v10_lattices(X23)|~l3_lattices(X23))|(~m1_subset_1(X24,u1_struct_0(X23))|k4_lattice3(X23,X24)=X24)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d3_lattice3])])])])).
cnf(c_0_32, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(k3_lattice3(k1_lattice3(esk1_0))))), inference(rw,[status(thm)],[c_0_27, c_0_17])).
cnf(c_0_33, plain, (r1_lattices(X1,X2,X3)|v2_struct_0(X1)|~r3_lattices(X1,X2,X3)|~v6_lattices(X1)|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k1_lattice3(esk1_0)))), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_23]), c_0_24]), c_0_25])]), c_0_26]), c_0_30])).
fof(c_0_35, plain, ![X22]:((~v2_struct_0(X22)|(v2_struct_0(X22)|~v10_lattices(X22))|~l3_lattices(X22))&((v4_lattices(X22)|(v2_struct_0(X22)|~v10_lattices(X22))|~l3_lattices(X22))&((v5_lattices(X22)|(v2_struct_0(X22)|~v10_lattices(X22))|~l3_lattices(X22))&((v6_lattices(X22)|(v2_struct_0(X22)|~v10_lattices(X22))|~l3_lattices(X22))&((v7_lattices(X22)|(v2_struct_0(X22)|~v10_lattices(X22))|~l3_lattices(X22))&((v8_lattices(X22)|(v2_struct_0(X22)|~v10_lattices(X22))|~l3_lattices(X22))&(v9_lattices(X22)|(v2_struct_0(X22)|~v10_lattices(X22))|~l3_lattices(X22)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
fof(c_0_36, plain, ![X11, X12, X13]:((~r3_lattices(X11,X12,X13)|r3_orders_2(k3_lattice3(X11),k4_lattice3(X11,X12),k4_lattice3(X11,X13))|~m1_subset_1(X13,u1_struct_0(X11))|~m1_subset_1(X12,u1_struct_0(X11))|(v2_struct_0(X11)|(~v10_lattices(X11)|~l3_lattices(X11))))&(~r3_orders_2(k3_lattice3(X11),k4_lattice3(X11,X12),k4_lattice3(X11,X13))|r3_lattices(X11,X12,X13)|~m1_subset_1(X13,u1_struct_0(X11))|~m1_subset_1(X12,u1_struct_0(X11))|(v2_struct_0(X11)|(~v10_lattices(X11)|~l3_lattices(X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_lattice3])])])])])).
cnf(c_0_37, plain, (v2_struct_0(X1)|k4_lattice3(X1,X2)=X2|~v10_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, negated_conjecture, (k5_lattice3(k1_lattice3(esk1_0),esk2_0)=esk2_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_32]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_39, negated_conjecture, (r1_lattices(k1_lattice3(esk1_0),X1,esk3_0)|~v9_lattices(k1_lattice3(esk1_0))|~v8_lattices(k1_lattice3(esk1_0))|~v6_lattices(k1_lattice3(esk1_0))|~r3_lattices(k1_lattice3(esk1_0),X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(k1_lattice3(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_24])]), c_0_26])).
cnf(c_0_40, plain, (v9_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41, plain, (r3_lattices(X1,X2,X3)|v2_struct_0(X1)|~r3_orders_2(k3_lattice3(X1),k4_lattice3(X1,X2),k4_lattice3(X1,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_42, negated_conjecture, (k4_lattice3(k1_lattice3(esk1_0),esk3_0)=esk3_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_34]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_43, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(k1_lattice3(esk1_0)))), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_32]), c_0_24]), c_0_25])]), c_0_26]), c_0_38])).
cnf(c_0_44, negated_conjecture, (r1_lattices(k1_lattice3(esk1_0),X1,esk3_0)|~v8_lattices(k1_lattice3(esk1_0))|~v6_lattices(k1_lattice3(esk1_0))|~r3_lattices(k1_lattice3(esk1_0),X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(k1_lattice3(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_45, plain, (v8_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_46, negated_conjecture, (r3_lattices(k1_lattice3(esk1_0),X1,esk3_0)|~r3_orders_2(k3_lattice3(k1_lattice3(esk1_0)),k4_lattice3(k1_lattice3(esk1_0),X1),esk3_0)|~m1_subset_1(X1,u1_struct_0(k1_lattice3(esk1_0)))), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_34]), c_0_24]), c_0_25])]), c_0_26]), c_0_42])).
cnf(c_0_47, negated_conjecture, (k4_lattice3(k1_lattice3(esk1_0),esk2_0)=esk2_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_43]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_48, negated_conjecture, (r3_orders_2(k3_yellow_1(esk1_0),esk2_0,esk3_0)|r1_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_49, plain, (r3_lattices(X1,X2,X3)|v2_struct_0(X1)|~r1_lattices(X1,X2,X3)|~v6_lattices(X1)|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_50, negated_conjecture, (r1_lattices(k1_lattice3(esk1_0),X1,esk3_0)|~v6_lattices(k1_lattice3(esk1_0))|~r3_lattices(k1_lattice3(esk1_0),X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(k1_lattice3(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_51, plain, (v6_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_52, negated_conjecture, (r3_lattices(k1_lattice3(esk1_0),esk2_0,esk3_0)|~r3_orders_2(k3_lattice3(k1_lattice3(esk1_0)),esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_43]), c_0_47])).
cnf(c_0_53, negated_conjecture, (r1_tarski(esk2_0,esk3_0)|r3_orders_2(k3_lattice3(k1_lattice3(esk1_0)),esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_48, c_0_17])).
cnf(c_0_54, negated_conjecture, (r3_lattices(k1_lattice3(esk1_0),X1,esk3_0)|~v9_lattices(k1_lattice3(esk1_0))|~v8_lattices(k1_lattice3(esk1_0))|~v6_lattices(k1_lattice3(esk1_0))|~r1_lattices(k1_lattice3(esk1_0),X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(k1_lattice3(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_34]), c_0_24])]), c_0_26])).
fof(c_0_55, plain, ![X14, X15, X16]:((~r1_lattices(k1_lattice3(X14),X15,X16)|r1_tarski(X15,X16)|~m1_subset_1(X16,u1_struct_0(k1_lattice3(X14)))|~m1_subset_1(X15,u1_struct_0(k1_lattice3(X14))))&(~r1_tarski(X15,X16)|r1_lattices(k1_lattice3(X14),X15,X16)|~m1_subset_1(X16,u1_struct_0(k1_lattice3(X14)))|~m1_subset_1(X15,u1_struct_0(k1_lattice3(X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_lattice3])])])])).
cnf(c_0_56, negated_conjecture, (r1_lattices(k1_lattice3(esk1_0),X1,esk3_0)|~r3_lattices(k1_lattice3(esk1_0),X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(k1_lattice3(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_57, negated_conjecture, (r3_lattices(k1_lattice3(esk1_0),esk2_0,esk3_0)|r1_tarski(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_58, negated_conjecture, (r3_lattices(k1_lattice3(esk1_0),X1,esk3_0)|~v8_lattices(k1_lattice3(esk1_0))|~v6_lattices(k1_lattice3(esk1_0))|~r1_lattices(k1_lattice3(esk1_0),X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(k1_lattice3(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_40]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_59, plain, (r1_lattices(k1_lattice3(X3),X1,X2)|~r1_tarski(X1,X2)|~m1_subset_1(X2,u1_struct_0(k1_lattice3(X3)))|~m1_subset_1(X1,u1_struct_0(k1_lattice3(X3)))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_60, plain, (r1_tarski(X2,X3)|~r1_lattices(k1_lattice3(X1),X2,X3)|~m1_subset_1(X3,u1_struct_0(k1_lattice3(X1)))|~m1_subset_1(X2,u1_struct_0(k1_lattice3(X1)))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_61, negated_conjecture, (r1_lattices(k1_lattice3(esk1_0),esk2_0,esk3_0)|r1_tarski(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_43])])).
cnf(c_0_62, negated_conjecture, (r3_lattices(k1_lattice3(esk1_0),X1,esk3_0)|~v6_lattices(k1_lattice3(esk1_0))|~r1_lattices(k1_lattice3(esk1_0),X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(k1_lattice3(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_45]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_63, negated_conjecture, (r1_lattices(k1_lattice3(esk1_0),X1,esk3_0)|~r1_tarski(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(k1_lattice3(esk1_0)))), inference(spm,[status(thm)],[c_0_59, c_0_34])).
cnf(c_0_64, negated_conjecture, (r1_tarski(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_34]), c_0_43])])).
cnf(c_0_65, negated_conjecture, (~r3_orders_2(k3_yellow_1(esk1_0),esk2_0,esk3_0)|~r1_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_66, plain, (r3_orders_2(k3_lattice3(X1),k4_lattice3(X1,X2),k4_lattice3(X1,X3))|v2_struct_0(X1)|~r3_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_67, negated_conjecture, (r3_lattices(k1_lattice3(esk1_0),X1,esk3_0)|~r1_lattices(k1_lattice3(esk1_0),X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(k1_lattice3(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_51]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_68, negated_conjecture, (r1_lattices(k1_lattice3(esk1_0),esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_43])])).
cnf(c_0_69, negated_conjecture, (~r1_tarski(esk2_0,esk3_0)|~r3_orders_2(k3_lattice3(k1_lattice3(esk1_0)),esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_65, c_0_17])).
cnf(c_0_70, negated_conjecture, (r3_orders_2(k3_lattice3(k1_lattice3(esk1_0)),k4_lattice3(k1_lattice3(esk1_0),X1),esk3_0)|~r3_lattices(k1_lattice3(esk1_0),X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(k1_lattice3(esk1_0)))), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_34]), c_0_24]), c_0_25])]), c_0_26]), c_0_42])).
cnf(c_0_71, negated_conjecture, (r3_lattices(k1_lattice3(esk1_0),esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_43])])).
cnf(c_0_72, negated_conjecture, (~r3_orders_2(k3_lattice3(k1_lattice3(esk1_0)),esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_64])])).
cnf(c_0_73, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_47]), c_0_43])]), c_0_72]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 74
# Proof object clause steps            : 49
# Proof object formula steps           : 25
# Proof object conjectures             : 36
# Proof object clause conjectures      : 33
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 24
# Proof object simplifying inferences  : 86
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 25
# Processed clauses                    : 84
# ...of these trivial                  : 2
# ...subsumed                          : 1
# ...remaining for further processing  : 81
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 5
# Generated clauses                    : 45
# ...of the previous two non-trivial   : 43
# Contextual simplify-reflections      : 0
# Paramodulations                      : 45
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 47
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 31
# Current number of unprocessed clauses: 7
# ...number of literals in the above   : 50
# Current number of archived formulas  : 0
# Current number of archived clauses   : 35
# Clause-clause subsumption calls (NU) : 388
# Rec. Clause-clause subsumption calls : 74
# Non-unit clause-clause subsumptions  : 7
# Unit Clause-clause subsumption calls : 17
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4052

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 29024 pages
