# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/06f492b3f6f441459ce8e3ac5726990a.E_input', t4_subset)).
fof(dt_u1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/06f492b3f6f441459ce8e3ac5726990a.E_input', dt_u1_pre_topc)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/06f492b3f6f441459ce8e3ac5726990a.E_input', t3_subset)).
fof(t94_zfmisc_1, axiom, ![X1, X2]:(![X3]:(r2_hidden(X3,X1)=>r1_tarski(X3,X2))=>r1_tarski(k3_tarski(X1),X2)), file('data/MPTP/proofs/06f492b3f6f441459ce8e3ac5726990a.E_input', t94_zfmisc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/06f492b3f6f441459ce8e3ac5726990a.E_input', d10_xboole_0)).
fof(l49_zfmisc_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>r1_tarski(X1,k3_tarski(X2))), file('data/MPTP/proofs/06f492b3f6f441459ce8e3ac5726990a.E_input', l49_zfmisc_1)).
fof(t24_yellow_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1)), file('data/MPTP/proofs/06f492b3f6f441459ce8e3ac5726990a.E_input', t24_yellow_1)).
fof(d1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1))))))))), file('data/MPTP/proofs/06f492b3f6f441459ce8e3ac5726990a.E_input', d1_pre_topc)).
fof(t14_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>(r2_hidden(k3_tarski(X1),X1)=>k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1))), file('data/MPTP/proofs/06f492b3f6f441459ce8e3ac5726990a.E_input', t14_yellow_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/06f492b3f6f441459ce8e3ac5726990a.E_input', d1_xboole_0)).
fof(c_0_10, plain, ![X4, X5, X6]:(~r2_hidden(X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(X6))|m1_subset_1(X4,X6)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_11, plain, ![X7]:(~l1_pre_topc(X7)|m1_subset_1(u1_pre_topc(X7),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X7))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_pre_topc])])).
fof(c_0_12, plain, ![X8, X9]:((~m1_subset_1(X8,k1_zfmisc_1(X9))|r1_tarski(X8,X9))&(~r1_tarski(X8,X9)|m1_subset_1(X8,k1_zfmisc_1(X9)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_13, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_15, plain, ![X10, X11]:((r2_hidden(esk1_2(X10,X11),X10)|r1_tarski(k3_tarski(X10),X11))&(~r1_tarski(esk1_2(X10,X11),X11)|r1_tarski(k3_tarski(X10),X11))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t94_zfmisc_1])])])])).
cnf(c_0_16, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~r2_hidden(X1,u1_pre_topc(X2))), inference(spm,[status(thm)],[c_0_13, c_0_14])).
fof(c_0_18, plain, ![X13, X14]:(((r1_tarski(X13,X14)|X13!=X14)&(r1_tarski(X14,X13)|X13!=X14))&(~r1_tarski(X13,X14)|~r1_tarski(X14,X13)|X13=X14)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_19, plain, ![X15, X16]:(~r2_hidden(X15,X16)|r1_tarski(X15,k3_tarski(X16))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l49_zfmisc_1])])).
cnf(c_0_20, plain, (r1_tarski(k3_tarski(X1),X2)|~r1_tarski(esk1_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (r1_tarski(X1,u1_struct_0(X2))|~l1_pre_topc(X2)|~r2_hidden(X1,u1_pre_topc(X2))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
fof(c_0_22, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1))), inference(assume_negation,[status(cth)],[t24_yellow_1])).
cnf(c_0_23, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (r1_tarski(k3_tarski(X1),u1_struct_0(X2))|~l1_pre_topc(X2)|~r2_hidden(esk1_2(X1,u1_struct_0(X2)),u1_pre_topc(X2))), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_26, plain, (r2_hidden(esk1_2(X1,X2),X1)|r1_tarski(k3_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_27, plain, ![X18, X19, X20, X21]:(((r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|~v2_pre_topc(X18)|~l1_pre_topc(X18))&((~m1_subset_1(X19,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X18))))|(~r1_tarski(X19,u1_pre_topc(X18))|r2_hidden(k5_setfam_1(u1_struct_0(X18),X19),u1_pre_topc(X18)))|~v2_pre_topc(X18)|~l1_pre_topc(X18))&(~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X18)))|(~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X18)))|(~r2_hidden(X20,u1_pre_topc(X18))|~r2_hidden(X21,u1_pre_topc(X18))|r2_hidden(k9_subset_1(u1_struct_0(X18),X20,X21),u1_pre_topc(X18))))|~v2_pre_topc(X18)|~l1_pre_topc(X18))))&(((m1_subset_1(esk4_1(X18),k1_zfmisc_1(u1_struct_0(X18)))|m1_subset_1(esk3_1(X18),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X18))))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18))&((m1_subset_1(esk5_1(X18),k1_zfmisc_1(u1_struct_0(X18)))|m1_subset_1(esk3_1(X18),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X18))))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18))&(((r2_hidden(esk4_1(X18),u1_pre_topc(X18))|m1_subset_1(esk3_1(X18),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X18))))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18))&(r2_hidden(esk5_1(X18),u1_pre_topc(X18))|m1_subset_1(esk3_1(X18),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X18))))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18)))&(~r2_hidden(k9_subset_1(u1_struct_0(X18),esk4_1(X18),esk5_1(X18)),u1_pre_topc(X18))|m1_subset_1(esk3_1(X18),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X18))))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18)))))&(((m1_subset_1(esk4_1(X18),k1_zfmisc_1(u1_struct_0(X18)))|r1_tarski(esk3_1(X18),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18))&((m1_subset_1(esk5_1(X18),k1_zfmisc_1(u1_struct_0(X18)))|r1_tarski(esk3_1(X18),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18))&(((r2_hidden(esk4_1(X18),u1_pre_topc(X18))|r1_tarski(esk3_1(X18),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18))&(r2_hidden(esk5_1(X18),u1_pre_topc(X18))|r1_tarski(esk3_1(X18),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18)))&(~r2_hidden(k9_subset_1(u1_struct_0(X18),esk4_1(X18),esk5_1(X18)),u1_pre_topc(X18))|r1_tarski(esk3_1(X18),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18)))))&((m1_subset_1(esk4_1(X18),k1_zfmisc_1(u1_struct_0(X18)))|~r2_hidden(k5_setfam_1(u1_struct_0(X18),esk3_1(X18)),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18))&((m1_subset_1(esk5_1(X18),k1_zfmisc_1(u1_struct_0(X18)))|~r2_hidden(k5_setfam_1(u1_struct_0(X18),esk3_1(X18)),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18))&(((r2_hidden(esk4_1(X18),u1_pre_topc(X18))|~r2_hidden(k5_setfam_1(u1_struct_0(X18),esk3_1(X18)),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18))&(r2_hidden(esk5_1(X18),u1_pre_topc(X18))|~r2_hidden(k5_setfam_1(u1_struct_0(X18),esk3_1(X18)),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18)))&(~r2_hidden(k9_subset_1(u1_struct_0(X18),esk4_1(X18),esk5_1(X18)),u1_pre_topc(X18))|~r2_hidden(k5_setfam_1(u1_struct_0(X18),esk3_1(X18)),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])).
fof(c_0_28, negated_conjecture, ((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&k4_yellow_0(k2_yellow_1(u1_pre_topc(esk2_0)))!=u1_struct_0(esk2_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_22])])])])).
fof(c_0_29, plain, ![X25]:(v1_xboole_0(X25)|(~r2_hidden(k3_tarski(X25),X25)|k4_yellow_0(k2_yellow_1(X25))=k3_tarski(X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_yellow_1])])])).
fof(c_0_30, plain, ![X26, X27, X28]:((~v1_xboole_0(X26)|~r2_hidden(X27,X26))&(r2_hidden(esk6_1(X28),X28)|v1_xboole_0(X28))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_31, plain, (k3_tarski(X1)=X2|~r1_tarski(k3_tarski(X1),X2)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_32, plain, (r1_tarski(k3_tarski(u1_pre_topc(X1)),u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_33, plain, (r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_34, negated_conjecture, (v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36, plain, (v1_xboole_0(X1)|k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, plain, (k3_tarski(u1_pre_topc(X1))=u1_struct_0(X1)|~l1_pre_topc(X1)|~r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_39, negated_conjecture, (r2_hidden(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])])).
cnf(c_0_40, plain, (k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(csr,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_41, negated_conjecture, (k3_tarski(u1_pre_topc(esk2_0))=u1_struct_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_35])])).
cnf(c_0_42, negated_conjecture, (k4_yellow_0(k2_yellow_1(u1_pre_topc(esk2_0)))!=u1_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_43, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_39])]), c_0_42]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 44
# Proof object clause steps            : 23
# Proof object formula steps           : 21
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 35
# Processed clauses                    : 167
# ...of these trivial                  : 0
# ...subsumed                          : 39
# ...remaining for further processing  : 128
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 0
# Generated clauses                    : 363
# ...of the previous two non-trivial   : 316
# Contextual simplify-reflections      : 1
# Paramodulations                      : 361
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 88
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 79
# Current number of unprocessed clauses: 217
# ...number of literals in the above   : 934
# Current number of archived formulas  : 0
# Current number of archived clauses   : 38
# Clause-clause subsumption calls (NU) : 2633
# Rec. Clause-clause subsumption calls : 891
# Non-unit clause-clause subsumptions  : 44
# Unit Clause-clause subsumption calls : 16
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8329

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.006 s
# Total time               : 0.014 s
# Maximum resident set size: 28676 pages
