# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t173_relat_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>(k10_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k2_relat_1(X2),X1))), file('data/MPTP/proofs/0748131e2d654196b74581709e17dd7f.E_input', t173_relat_1)).
fof(l24_zfmisc_1, axiom, ![X1, X2]:~((r1_xboole_0(k1_tarski(X1),X2)&r2_hidden(X1,X2))), file('data/MPTP/proofs/0748131e2d654196b74581709e17dd7f.E_input', l24_zfmisc_1)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/0748131e2d654196b74581709e17dd7f.E_input', t65_xboole_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/0748131e2d654196b74581709e17dd7f.E_input', d7_xboole_0)).
fof(t166_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k10_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k2_relat_1(X3))&(r2_hidden(k4_tarski(X1,X4),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/0748131e2d654196b74581709e17dd7f.E_input', t166_relat_1)).
fof(t168_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k10_relat_1(X2,X1)=k10_relat_1(X2,k3_xboole_0(k2_relat_1(X2),X1))), file('data/MPTP/proofs/0748131e2d654196b74581709e17dd7f.E_input', t168_relat_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/0748131e2d654196b74581709e17dd7f.E_input', t7_xboole_0)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/0748131e2d654196b74581709e17dd7f.E_input', d4_xboole_0)).
fof(d14_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(X3=k10_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(k4_tarski(X4,X5),X1)&r2_hidden(X5,X2))))), file('data/MPTP/proofs/0748131e2d654196b74581709e17dd7f.E_input', d14_relat_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/0748131e2d654196b74581709e17dd7f.E_input', d5_relat_1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/0748131e2d654196b74581709e17dd7f.E_input', t4_xboole_0)).
fof(c_0_11, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>(k10_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k2_relat_1(X2),X1)))), inference(assume_negation,[status(cth)],[t173_relat_1])).
fof(c_0_12, plain, ![X8, X9]:(~r1_xboole_0(k1_tarski(X8),X9)|~r2_hidden(X8,X9)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l24_zfmisc_1])])).
fof(c_0_13, plain, ![X10]:r1_xboole_0(X10,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
fof(c_0_14, plain, ![X11, X12]:((~r1_xboole_0(X11,X12)|k3_xboole_0(X11,X12)=k1_xboole_0)&(k3_xboole_0(X11,X12)!=k1_xboole_0|r1_xboole_0(X11,X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
fof(c_0_15, negated_conjecture, (v1_relat_1(esk2_0)&((k10_relat_1(esk2_0,esk1_0)!=k1_xboole_0|~r1_xboole_0(k2_relat_1(esk2_0),esk1_0))&(k10_relat_1(esk2_0,esk1_0)=k1_xboole_0|r1_xboole_0(k2_relat_1(esk2_0),esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
cnf(c_0_16, plain, (~r1_xboole_0(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_18, plain, ![X13, X14, X15, X17]:(((r2_hidden(esk3_3(X13,X14,X15),k2_relat_1(X15))|~r2_hidden(X13,k10_relat_1(X15,X14))|~v1_relat_1(X15))&((r2_hidden(k4_tarski(X13,esk3_3(X13,X14,X15)),X15)|~r2_hidden(X13,k10_relat_1(X15,X14))|~v1_relat_1(X15))&(r2_hidden(esk3_3(X13,X14,X15),X14)|~r2_hidden(X13,k10_relat_1(X15,X14))|~v1_relat_1(X15))))&(~r2_hidden(X17,k2_relat_1(X15))|(~r2_hidden(k4_tarski(X13,X17),X15)|~r2_hidden(X17,X14))|r2_hidden(X13,k10_relat_1(X15,X14))|~v1_relat_1(X15))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t166_relat_1])])])])])).
fof(c_0_19, plain, ![X18, X19]:(~v1_relat_1(X19)|k10_relat_1(X19,X18)=k10_relat_1(X19,k3_xboole_0(k2_relat_1(X19),X18))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t168_relat_1])])).
cnf(c_0_20, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)=k1_xboole_0|r1_xboole_0(k2_relat_1(esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_23, plain, (r2_hidden(esk3_3(X1,X2,X3),X2)|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_24, plain, ![X20]:(X20=k1_xboole_0|r2_hidden(esk4_1(X20),X20)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_25, plain, (k10_relat_1(X1,X2)=k10_relat_1(X1,k3_xboole_0(k2_relat_1(X1),X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, negated_conjecture, (k3_xboole_0(k2_relat_1(esk2_0),esk1_0)=k1_xboole_0|k10_relat_1(esk2_0,esk1_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_28, plain, (~r2_hidden(X1,k10_relat_1(X2,k1_xboole_0))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_29, plain, (X1=k1_xboole_0|r2_hidden(esk4_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)=k10_relat_1(esk2_0,k1_xboole_0)|k10_relat_1(esk2_0,esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])])).
cnf(c_0_31, plain, (k10_relat_1(X1,k1_xboole_0)=k1_xboole_0|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
fof(c_0_32, plain, ![X22, X23, X24, X25, X26, X27, X28, X29]:((((r2_hidden(X25,X22)|~r2_hidden(X25,X24)|X24!=k3_xboole_0(X22,X23))&(r2_hidden(X25,X23)|~r2_hidden(X25,X24)|X24!=k3_xboole_0(X22,X23)))&(~r2_hidden(X26,X22)|~r2_hidden(X26,X23)|r2_hidden(X26,X24)|X24!=k3_xboole_0(X22,X23)))&((~r2_hidden(esk5_3(X27,X28,X29),X29)|(~r2_hidden(esk5_3(X27,X28,X29),X27)|~r2_hidden(esk5_3(X27,X28,X29),X28))|X29=k3_xboole_0(X27,X28))&((r2_hidden(esk5_3(X27,X28,X29),X27)|r2_hidden(esk5_3(X27,X28,X29),X29)|X29=k3_xboole_0(X27,X28))&(r2_hidden(esk5_3(X27,X28,X29),X28)|r2_hidden(esk5_3(X27,X28,X29),X29)|X29=k3_xboole_0(X27,X28))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_33, plain, ![X31, X32, X33, X34, X36, X37, X38, X39, X41]:((((r2_hidden(k4_tarski(X34,esk6_4(X31,X32,X33,X34)),X31)|~r2_hidden(X34,X33)|X33!=k10_relat_1(X31,X32)|~v1_relat_1(X31))&(r2_hidden(esk6_4(X31,X32,X33,X34),X32)|~r2_hidden(X34,X33)|X33!=k10_relat_1(X31,X32)|~v1_relat_1(X31)))&(~r2_hidden(k4_tarski(X36,X37),X31)|~r2_hidden(X37,X32)|r2_hidden(X36,X33)|X33!=k10_relat_1(X31,X32)|~v1_relat_1(X31)))&((~r2_hidden(esk7_3(X31,X38,X39),X39)|(~r2_hidden(k4_tarski(esk7_3(X31,X38,X39),X41),X31)|~r2_hidden(X41,X38))|X39=k10_relat_1(X31,X38)|~v1_relat_1(X31))&((r2_hidden(k4_tarski(esk7_3(X31,X38,X39),esk8_3(X31,X38,X39)),X31)|r2_hidden(esk7_3(X31,X38,X39),X39)|X39=k10_relat_1(X31,X38)|~v1_relat_1(X31))&(r2_hidden(esk8_3(X31,X38,X39),X38)|r2_hidden(esk7_3(X31,X38,X39),X39)|X39=k10_relat_1(X31,X38)|~v1_relat_1(X31))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d14_relat_1])])])])])])).
cnf(c_0_34, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)=k1_xboole_0|k10_relat_1(esk2_0,k1_xboole_0)!=k1_xboole_0), inference(ef,[status(thm)],[c_0_30])).
cnf(c_0_35, negated_conjecture, (k10_relat_1(esk2_0,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_31, c_0_27])).
cnf(c_0_36, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_37, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_20, c_0_17])).
cnf(c_0_38, plain, (r2_hidden(X1,X5)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X2,X4)|X5!=k10_relat_1(X3,X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_39, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_35])])).
cnf(c_0_40, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_22])).
cnf(c_0_41, negated_conjecture, (X1!=k1_xboole_0|~r2_hidden(k4_tarski(X2,X3),esk2_0)|~r2_hidden(X3,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_27])]), c_0_40])).
fof(c_0_42, plain, ![X43, X44, X45, X47, X48, X49, X50, X52]:(((~r2_hidden(X45,X44)|r2_hidden(k4_tarski(esk9_3(X43,X44,X45),X45),X43)|X44!=k2_relat_1(X43))&(~r2_hidden(k4_tarski(X48,X47),X43)|r2_hidden(X47,X44)|X44!=k2_relat_1(X43)))&((~r2_hidden(esk10_2(X49,X50),X50)|~r2_hidden(k4_tarski(X52,esk10_2(X49,X50)),X49)|X50=k2_relat_1(X49))&(r2_hidden(esk10_2(X49,X50),X50)|r2_hidden(k4_tarski(esk11_2(X49,X50),esk10_2(X49,X50)),X49)|X50=k2_relat_1(X49)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
fof(c_0_43, plain, ![X54, X55, X57, X58, X59]:((r1_xboole_0(X54,X55)|r2_hidden(esk12_2(X54,X55),k3_xboole_0(X54,X55)))&(~r2_hidden(X59,k3_xboole_0(X57,X58))|~r1_xboole_0(X57,X58))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
cnf(c_0_44, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X4!=k3_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_45, negated_conjecture, (~r2_hidden(k4_tarski(X1,X2),esk2_0)|~r2_hidden(X2,esk1_0)), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_46, plain, (r2_hidden(k4_tarski(esk9_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_47, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_48, plain, (r2_hidden(X1,k3_xboole_0(X2,X3))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_49, negated_conjecture, (X1!=k2_relat_1(esk2_0)|~r2_hidden(X2,esk1_0)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_50, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_36])).
cnf(c_0_51, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk12_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_52, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_53, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_54, plain, (r2_hidden(esk5_3(X1,X2,X3),X1)|r2_hidden(esk5_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_55, plain, (r2_hidden(esk5_3(X1,X2,X3),X2)|r2_hidden(esk5_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_56, negated_conjecture, (~r2_hidden(X1,k2_relat_1(esk2_0))|~r2_hidden(X1,esk1_0)), inference(er,[status(thm)],[c_0_49])).
cnf(c_0_57, plain, (r2_hidden(esk12_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_58, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_52])).
cnf(c_0_59, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)!=k1_xboole_0|~r1_xboole_0(k2_relat_1(esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_60, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_61, plain, (X1=k3_xboole_0(X2,X3)|r2_hidden(esk5_3(X2,X3,X1),X1)|~r2_hidden(esk5_3(X2,X3,X1),X4)|~r1_xboole_0(X4,X2)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_62, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|r2_hidden(esk5_3(X1,X2,k1_xboole_0),X2)), inference(spm,[status(thm)],[c_0_22, c_0_55])).
cnf(c_0_63, negated_conjecture, (r1_xboole_0(X1,k2_relat_1(esk2_0))|~r2_hidden(esk12_2(X1,k2_relat_1(esk2_0)),esk1_0)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_64, plain, (r2_hidden(esk12_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_58, c_0_51])).
cnf(c_0_65, negated_conjecture, (k3_xboole_0(k2_relat_1(esk2_0),esk1_0)!=k1_xboole_0|k10_relat_1(esk2_0,esk1_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_66, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_22])).
cnf(c_0_67, negated_conjecture, (r1_xboole_0(esk1_0,k2_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_68, negated_conjecture, (k3_xboole_0(k2_relat_1(esk2_0),esk1_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_39])])).
cnf(c_0_69, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 70
# Proof object clause steps            : 47
# Proof object formula steps           : 23
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 26
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 31
# Processed clauses                    : 3033
# ...of these trivial                  : 18
# ...subsumed                          : 2502
# ...remaining for further processing  : 513
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 30
# Backward-rewritten                   : 14
# Generated clauses                    : 11813
# ...of the previous two non-trivial   : 10440
# Contextual simplify-reflections      : 1
# Paramodulations                      : 11646
# Factorizations                       : 23
# Equation resolutions                 : 144
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 438
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 421
# Current number of unprocessed clauses: 7417
# ...number of literals in the above   : 30747
# Current number of archived formulas  : 0
# Current number of archived clauses   : 75
# Clause-clause subsumption calls (NU) : 77022
# Rec. Clause-clause subsumption calls : 29640
# Non-unit clause-clause subsumptions  : 1950
# Unit Clause-clause subsumption calls : 303
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 7
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 139005

# -------------------------------------------------
# User time                : 0.119 s
# System time              : 0.003 s
# Total time               : 0.122 s
# Maximum resident set size: 28792 pages
