# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t130_funct_2, conjecture, ![X1, X2, X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>(k1_relset_1(X1,X2,X3)=X1=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))), file('data/MPTP/proofs/07a0f8317d894bdbaa09cfcb8d7f2a9b.E_input', t130_funct_2)).
fof(t69_xboole_1, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>~((r1_tarski(X2,X1)&r1_xboole_0(X2,X1)))), file('data/MPTP/proofs/07a0f8317d894bdbaa09cfcb8d7f2a9b.E_input', t69_xboole_1)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/07a0f8317d894bdbaa09cfcb8d7f2a9b.E_input', t65_xboole_1)).
fof(cc4_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/07a0f8317d894bdbaa09cfcb8d7f2a9b.E_input', cc4_relset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/07a0f8317d894bdbaa09cfcb8d7f2a9b.E_input', t3_subset)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/07a0f8317d894bdbaa09cfcb8d7f2a9b.E_input', l13_xboole_0)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(((X2=k1_xboole_0=>X1=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X2,X3)))&(X2=k1_xboole_0=>(X1=k1_xboole_0|(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0))))), file('data/MPTP/proofs/07a0f8317d894bdbaa09cfcb8d7f2a9b.E_input', d1_funct_2)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/07a0f8317d894bdbaa09cfcb8d7f2a9b.E_input', t2_xboole_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/07a0f8317d894bdbaa09cfcb8d7f2a9b.E_input', t113_zfmisc_1)).
fof(rc2_funct_2, axiom, ![X1]:?[X2]:(m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))&(v1_relat_1(X2)&(v4_relat_1(X2,X1)&(v5_relat_1(X2,X1)&(v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&v3_funct_2(X2,X1,X1))))))), file('data/MPTP/proofs/07a0f8317d894bdbaa09cfcb8d7f2a9b.E_input', rc2_funct_2)).
fof(c_0_10, negated_conjecture, ~(![X1, X2, X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>(k1_relset_1(X1,X2,X3)=X1=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))))), inference(assume_negation,[status(cth)],[t130_funct_2])).
fof(c_0_11, plain, ![X17, X18]:(v1_xboole_0(X18)|(~r1_tarski(X18,X17)|~r1_xboole_0(X18,X17))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t69_xboole_1])])])).
fof(c_0_12, plain, ![X7]:r1_xboole_0(X7,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
fof(c_0_13, plain, ![X19, X20, X21]:(~v1_xboole_0(X19)|(~m1_subset_1(X21,k1_zfmisc_1(k2_zfmisc_1(X20,X19)))|v1_xboole_0(X21))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])).
fof(c_0_14, negated_conjecture, ((v1_funct_1(esk3_0)&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0))))&(k1_relset_1(esk1_0,esk2_0,esk3_0)=esk1_0&(~v1_funct_1(esk3_0)|(~v1_funct_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
cnf(c_0_15, plain, (v1_xboole_0(X1)|~r1_tarski(X1,X2)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_17, plain, ![X22, X23]:((~m1_subset_1(X22,k1_zfmisc_1(X23))|r1_tarski(X22,X23))&(~r1_tarski(X22,X23)|m1_subset_1(X22,k1_zfmisc_1(X23)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_18, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (v1_xboole_0(X1)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_21, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_22, plain, ![X8]:(~v1_xboole_0(X8)|X8=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_23, negated_conjecture, (v1_xboole_0(esk3_0)|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_24, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_20, c_0_21])).
fof(c_0_25, plain, ![X9, X10, X11]:((((~v1_funct_2(X11,X9,X10)|X9=k1_relset_1(X9,X10,X11)|X10=k1_xboole_0|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X9,X10))))&(X9!=k1_relset_1(X9,X10,X11)|v1_funct_2(X11,X9,X10)|X10=k1_xboole_0|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))))&((~v1_funct_2(X11,X9,X10)|X9=k1_relset_1(X9,X10,X11)|X9!=k1_xboole_0|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X9,X10))))&(X9!=k1_relset_1(X9,X10,X11)|v1_funct_2(X11,X9,X10)|X9!=k1_xboole_0|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X9,X10))))))&((~v1_funct_2(X11,X9,X10)|X11=k1_xboole_0|X9=k1_xboole_0|X10!=k1_xboole_0|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X9,X10))))&(X11!=k1_xboole_0|v1_funct_2(X11,X9,X10)|X9=k1_xboole_0|X10!=k1_xboole_0|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
cnf(c_0_26, negated_conjecture, (~v1_funct_1(esk3_0)|~v1_funct_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_27, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_28, plain, ![X14]:r1_tarski(k1_xboole_0,X14), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_29, plain, ![X12, X13]:((k2_zfmisc_1(X12,X13)!=k1_xboole_0|(X12=k1_xboole_0|X13=k1_xboole_0))&((X12!=k1_xboole_0|k2_zfmisc_1(X12,X13)=k1_xboole_0)&(X13!=k1_xboole_0|k2_zfmisc_1(X12,X13)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_30, plain, ![X15]:(m1_subset_1(esk4_1(X15),k1_zfmisc_1(k2_zfmisc_1(X15,X15)))&(v1_relat_1(esk4_1(X15))&(v4_relat_1(esk4_1(X15),X15)&(v5_relat_1(esk4_1(X15),X15)&(v1_funct_1(esk4_1(X15))&(v1_funct_2(esk4_1(X15),X15,X15)&v3_funct_2(esk4_1(X15),X15,X15))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc2_funct_2])])).
cnf(c_0_31, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32, negated_conjecture, (v1_xboole_0(esk3_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_33, plain, (v1_funct_2(X3,X1,X2)|X2=k1_xboole_0|X1!=k1_relset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, negated_conjecture, (k1_relset_1(esk1_0,esk2_0,esk3_0)=esk1_0), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_35, negated_conjecture, (~v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_27]), c_0_19])])).
cnf(c_0_36, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_37, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40, negated_conjecture, (esk3_0=k1_xboole_0|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_41, negated_conjecture, (esk2_0=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_19]), c_0_34])]), c_0_35])).
cnf(c_0_42, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_43, plain, (v1_funct_2(X1,X2,X3)|X2=k1_xboole_0|X1!=k1_xboole_0|X3!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_44, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_45, plain, (v1_xboole_0(esk4_1(X1))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_18, c_0_39])).
cnf(c_0_46, negated_conjecture, (esk3_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_41]), c_0_42])])).
cnf(c_0_47, plain, (X1=k1_xboole_0|v1_funct_2(k1_xboole_0,X1,k1_xboole_0)|~m1_subset_1(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))), inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_43])]), c_0_44])).
cnf(c_0_48, plain, (esk4_1(X1)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_31, c_0_45])).
cnf(c_0_49, plain, (v1_xboole_0(k1_xboole_0)), inference(spm,[status(thm)],[c_0_20, c_0_37])).
cnf(c_0_50, negated_conjecture, (~v1_funct_2(k1_xboole_0,esk1_0,k1_xboole_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_41]), c_0_46])).
cnf(c_0_51, plain, (X1=k1_xboole_0|v1_funct_2(k1_xboole_0,X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_42])])).
cnf(c_0_52, plain, (v1_funct_2(esk4_1(X1),X1,X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_53, plain, (esk4_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_54, negated_conjecture, (esk1_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_55, plain, (v1_funct_2(k1_xboole_0,k1_xboole_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_56, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_54]), c_0_55])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 57
# Proof object clause steps            : 36
# Proof object formula steps           : 21
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 20
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 93
# ...of these trivial                  : 4
# ...subsumed                          : 2
# ...remaining for further processing  : 87
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 10
# Generated clauses                    : 58
# ...of the previous two non-trivial   : 47
# Contextual simplify-reflections      : 0
# Paramodulations                      : 52
# Factorizations                       : 0
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 44
#    Positive orientable unit clauses  : 23
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 21
# Current number of unprocessed clauses: 7
# ...number of literals in the above   : 21
# Current number of archived formulas  : 0
# Current number of archived clauses   : 37
# Clause-clause subsumption calls (NU) : 164
# Rec. Clause-clause subsumption calls : 104
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 8
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1981

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.003 s
# Total time               : 0.010 s
# Maximum resident set size: 28648 pages
