# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/081ae716558842c4acd12255e1cadbd2.E_input', d5_xboole_0)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/081ae716558842c4acd12255e1cadbd2.E_input', t100_xboole_1)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/081ae716558842c4acd12255e1cadbd2.E_input', d4_xboole_0)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/081ae716558842c4acd12255e1cadbd2.E_input', t3_boole)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/081ae716558842c4acd12255e1cadbd2.E_input', t7_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/081ae716558842c4acd12255e1cadbd2.E_input', commutativity_k3_xboole_0)).
fof(t65_zfmisc_1, conjecture, ![X1, X2]:(k4_xboole_0(X1,k1_tarski(X2))=X1<=>~(r2_hidden(X2,X1))), file('data/MPTP/proofs/081ae716558842c4acd12255e1cadbd2.E_input', t65_zfmisc_1)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/081ae716558842c4acd12255e1cadbd2.E_input', t83_xboole_1)).
fof(l27_zfmisc_1, axiom, ![X1, X2]:(~(r2_hidden(X1,X2))=>r1_xboole_0(k1_tarski(X1),X2)), file('data/MPTP/proofs/081ae716558842c4acd12255e1cadbd2.E_input', l27_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/081ae716558842c4acd12255e1cadbd2.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/081ae716558842c4acd12255e1cadbd2.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/081ae716558842c4acd12255e1cadbd2.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/081ae716558842c4acd12255e1cadbd2.E_input', t72_enumset1)).
fof(d1_enumset1, axiom, ![X1, X2, X3, X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3)))))), file('data/MPTP/proofs/081ae716558842c4acd12255e1cadbd2.E_input', d1_enumset1)).
fof(c_0_14, plain, ![X47, X48, X49, X50, X51, X52, X53, X54]:((((r2_hidden(X50,X47)|~r2_hidden(X50,X49)|X49!=k4_xboole_0(X47,X48))&(~r2_hidden(X50,X48)|~r2_hidden(X50,X49)|X49!=k4_xboole_0(X47,X48)))&(~r2_hidden(X51,X47)|r2_hidden(X51,X48)|r2_hidden(X51,X49)|X49!=k4_xboole_0(X47,X48)))&((~r2_hidden(esk6_3(X52,X53,X54),X54)|(~r2_hidden(esk6_3(X52,X53,X54),X52)|r2_hidden(esk6_3(X52,X53,X54),X53))|X54=k4_xboole_0(X52,X53))&((r2_hidden(esk6_3(X52,X53,X54),X52)|r2_hidden(esk6_3(X52,X53,X54),X54)|X54=k4_xboole_0(X52,X53))&(~r2_hidden(esk6_3(X52,X53,X54),X53)|r2_hidden(esk6_3(X52,X53,X54),X54)|X54=k4_xboole_0(X52,X53))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_15, plain, ![X8, X9]:k4_xboole_0(X8,X9)=k5_xboole_0(X8,k3_xboole_0(X8,X9)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_16, plain, ![X37, X38, X39, X40, X41, X42, X43, X44]:((((r2_hidden(X40,X37)|~r2_hidden(X40,X39)|X39!=k3_xboole_0(X37,X38))&(r2_hidden(X40,X38)|~r2_hidden(X40,X39)|X39!=k3_xboole_0(X37,X38)))&(~r2_hidden(X41,X37)|~r2_hidden(X41,X38)|r2_hidden(X41,X39)|X39!=k3_xboole_0(X37,X38)))&((~r2_hidden(esk5_3(X42,X43,X44),X44)|(~r2_hidden(esk5_3(X42,X43,X44),X42)|~r2_hidden(esk5_3(X42,X43,X44),X43))|X44=k3_xboole_0(X42,X43))&((r2_hidden(esk5_3(X42,X43,X44),X42)|r2_hidden(esk5_3(X42,X43,X44),X44)|X44=k3_xboole_0(X42,X43))&(r2_hidden(esk5_3(X42,X43,X44),X43)|r2_hidden(esk5_3(X42,X43,X44),X44)|X44=k3_xboole_0(X42,X43))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_17, plain, ![X17]:k4_xboole_0(X17,k1_xboole_0)=X17, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_18, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_21, plain, ![X35]:(X35=k1_xboole_0|r2_hidden(esk4_1(X35),X35)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_22, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_23, plain, ![X20, X21]:k3_xboole_0(X20,X21)=k3_xboole_0(X21,X20), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_24, negated_conjecture, ~(![X1, X2]:(k4_xboole_0(X1,k1_tarski(X2))=X1<=>~(r2_hidden(X2,X1)))), inference(assume_negation,[status(cth)],[t65_zfmisc_1])).
cnf(c_0_25, plain, (X3!=k5_xboole_0(X4,k3_xboole_0(X4,X2))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_26, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (X1=k1_xboole_0|r2_hidden(esk4_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_29, plain, ![X22, X23]:((~r1_xboole_0(X22,X23)|k4_xboole_0(X22,X23)=X22)&(k4_xboole_0(X22,X23)!=X22|r1_xboole_0(X22,X23))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
fof(c_0_30, plain, ![X59, X60]:(r2_hidden(X59,X60)|r1_xboole_0(k1_tarski(X59),X60)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l27_zfmisc_1])])])).
fof(c_0_31, plain, ![X46]:k2_tarski(X46,X46)=k1_tarski(X46), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_32, plain, ![X18, X19]:k1_enumset1(X18,X18,X19)=k2_tarski(X18,X19), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_33, plain, ![X14, X15, X16]:k2_enumset1(X14,X14,X15,X16)=k1_enumset1(X14,X15,X16), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_34, plain, ![X10, X11, X12, X13]:k3_enumset1(X10,X10,X11,X12,X13)=k2_enumset1(X10,X11,X12,X13), inference(variable_rename,[status(thm)],[t72_enumset1])).
cnf(c_0_35, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[c_0_22, c_0_19])).
cnf(c_0_36, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_37, negated_conjecture, ((k4_xboole_0(esk1_0,k1_tarski(esk2_0))!=esk1_0|r2_hidden(esk2_0,esk1_0))&(k4_xboole_0(esk1_0,k1_tarski(esk2_0))=esk1_0|~r2_hidden(esk2_0,esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_24])])])])).
fof(c_0_38, plain, ![X24, X25, X26, X27, X28, X29, X30, X31, X32, X33]:(((~r2_hidden(X28,X27)|(X28=X24|(X28=X25|X28=X26))|X27!=k1_enumset1(X24,X25,X26))&((X29!=X24|r2_hidden(X29,X27)|X27!=k1_enumset1(X24,X25,X26))&((X29!=X25|r2_hidden(X29,X27)|X27!=k1_enumset1(X24,X25,X26))&(X29!=X26|r2_hidden(X29,X27)|X27!=k1_enumset1(X24,X25,X26)))))&(((esk3_4(X30,X31,X32,X33)!=X30|~r2_hidden(esk3_4(X30,X31,X32,X33),X33)|X33=k1_enumset1(X30,X31,X32))&((esk3_4(X30,X31,X32,X33)!=X31|~r2_hidden(esk3_4(X30,X31,X32,X33),X33)|X33=k1_enumset1(X30,X31,X32))&(esk3_4(X30,X31,X32,X33)!=X32|~r2_hidden(esk3_4(X30,X31,X32,X33),X33)|X33=k1_enumset1(X30,X31,X32))))&(r2_hidden(esk3_4(X30,X31,X32,X33),X33)|(esk3_4(X30,X31,X32,X33)=X30|(esk3_4(X30,X31,X32,X33)=X31|esk3_4(X30,X31,X32,X33)=X32))|X33=k1_enumset1(X30,X31,X32)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])).
cnf(c_0_39, plain, (~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_25])).
cnf(c_0_40, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|r2_hidden(esk4_1(k3_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_41, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_28])).
cnf(c_0_42, plain, (k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_43, plain, (r2_hidden(X1,X2)|r1_xboole_0(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_44, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_45, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_46, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_47, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_48, plain, (k5_xboole_0(X1,k3_xboole_0(k1_xboole_0,X1))=X1), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_49, negated_conjecture, (k4_xboole_0(esk1_0,k1_tarski(esk2_0))=esk1_0|~r2_hidden(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_50, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X5,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_51, plain, (k3_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))=k1_xboole_0|~r2_hidden(esk4_1(k3_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))),X3)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_52, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|r2_hidden(esk4_1(k3_xboole_0(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_41, c_0_27])).
cnf(c_0_53, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=X1|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_42, c_0_19])).
cnf(c_0_54, plain, (r2_hidden(X1,X2)|r1_xboole_0(k3_enumset1(X1,X1,X1,X1,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_46]), c_0_47])).
cnf(c_0_55, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_39, c_0_48])).
cnf(c_0_56, negated_conjecture, (k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))=esk1_0|~r2_hidden(esk2_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_44]), c_0_45]), c_0_19]), c_0_46]), c_0_47])).
cnf(c_0_57, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k3_enumset1(X4,X4,X4,X5,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_46]), c_0_47])).
cnf(c_0_58, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|k4_xboole_0(esk1_0,k1_tarski(esk2_0))!=esk1_0), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_59, plain, (k3_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_60, plain, (k5_xboole_0(k3_enumset1(X1,X1,X1,X1,X1),k3_xboole_0(k3_enumset1(X1,X1,X1,X1,X1),X2))=k3_enumset1(X1,X1,X1,X1,X1)|r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_61, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_55, c_0_40])).
cnf(c_0_62, negated_conjecture, (~r2_hidden(X1,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))|~r2_hidden(esk2_0,esk1_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_39, c_0_56])).
cnf(c_0_63, plain, (r2_hidden(X1,k3_enumset1(X2,X2,X2,X3,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_57])])).
cnf(c_0_64, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))!=esk1_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_44]), c_0_45]), c_0_19]), c_0_46]), c_0_47])).
cnf(c_0_65, plain, (k3_xboole_0(X1,k3_enumset1(X2,X2,X2,X2,X2))=k1_xboole_0|r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_66, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_35, c_0_61])).
cnf(c_0_67, negated_conjecture, (~r2_hidden(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_68, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_66])]), c_0_67]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 69
# Proof object clause steps            : 40
# Proof object formula steps           : 29
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 28
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 37
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 32
# Processed clauses                    : 1460
# ...of these trivial                  : 458
# ...subsumed                          : 667
# ...remaining for further processing  : 335
# Other redundant clauses eliminated   : 22
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 21
# Generated clauses                    : 13514
# ...of the previous two non-trivial   : 8905
# Contextual simplify-reflections      : 2
# Paramodulations                      : 13445
# Factorizations                       : 50
# Equation resolutions                 : 22
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 267
#    Positive orientable unit clauses  : 94
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 170
# Current number of unprocessed clauses: 7461
# ...number of literals in the above   : 21629
# Current number of archived formulas  : 0
# Current number of archived clauses   : 61
# Clause-clause subsumption calls (NU) : 11535
# Rec. Clause-clause subsumption calls : 8681
# Non-unit clause-clause subsumptions  : 398
# Unit Clause-clause subsumption calls : 2380
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 315
# BW rewrite match successes           : 27
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 180192

# -------------------------------------------------
# User time                : 0.103 s
# System time              : 0.000 s
# Total time               : 0.103 s
# Maximum resident set size: 28716 pages
