# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_waybel_7, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v1_subset_1(X2,u1_struct_0(X1))<=>~(r2_hidden(k3_yellow_0(X1),X2))))), file('data/MPTP/proofs/0831cc725dfc4c8a9cecc959bd44a9f5.E_input', t8_waybel_7)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/0831cc725dfc4c8a9cecc959bd44a9f5.E_input', dt_k1_yellow_0)).
fof(t30_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4)))))&((r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))=>(X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3)))))), file('data/MPTP/proofs/0831cc725dfc4c8a9cecc959bd44a9f5.E_input', t30_yellow_0)).
fof(t49_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(![X3]:(m1_subset_1(X3,X1)=>r2_hidden(X3,X2))=>X1=X2)), file('data/MPTP/proofs/0831cc725dfc4c8a9cecc959bd44a9f5.E_input', t49_subset_1)).
fof(d20_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v13_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r1_orders_2(X1,X3,X4))=>r2_hidden(X4,X2))))))), file('data/MPTP/proofs/0831cc725dfc4c8a9cecc959bd44a9f5.E_input', d20_waybel_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/0831cc725dfc4c8a9cecc959bd44a9f5.E_input', t4_subset)).
fof(t6_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_lattice3(X1,k1_xboole_0,X2)&r1_lattice3(X1,k1_xboole_0,X2)))), file('data/MPTP/proofs/0831cc725dfc4c8a9cecc959bd44a9f5.E_input', t6_yellow_0)).
fof(t42_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))=>(r1_yellow_0(X1,k1_xboole_0)&r2_yellow_0(X1,u1_struct_0(X1)))), file('data/MPTP/proofs/0831cc725dfc4c8a9cecc959bd44a9f5.E_input', t42_yellow_0)).
fof(d11_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0)), file('data/MPTP/proofs/0831cc725dfc4c8a9cecc959bd44a9f5.E_input', d11_yellow_0)).
fof(d7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(v1_subset_1(X2,X1)<=>~(X2=X1))), file('data/MPTP/proofs/0831cc725dfc4c8a9cecc959bd44a9f5.E_input', d7_subset_1)).
fof(rc3_subset_1, axiom, ![X1]:?[X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))&~(v1_subset_1(X2,X1))), file('data/MPTP/proofs/0831cc725dfc4c8a9cecc959bd44a9f5.E_input', rc3_subset_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/0831cc725dfc4c8a9cecc959bd44a9f5.E_input', t2_subset)).
fof(c_0_12, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v1_subset_1(X2,u1_struct_0(X1))<=>~(r2_hidden(k3_yellow_0(X1),X2)))))), inference(assume_negation,[status(cth)],[t8_waybel_7])).
fof(c_0_13, plain, ![X7, X8]:(~l1_orders_2(X7)|m1_subset_1(k1_yellow_0(X7,X8),u1_struct_0(X7))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
fof(c_0_14, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_yellow_0(esk1_0)&l1_orders_2(esk1_0))))))&((~v1_xboole_0(esk2_0)&(v2_waybel_0(esk2_0,esk1_0)&(v13_waybel_0(esk2_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))))&((~v1_subset_1(esk2_0,u1_struct_0(esk1_0))|r2_hidden(k3_yellow_0(esk1_0),esk2_0))&(v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k3_yellow_0(esk1_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])).
fof(c_0_15, plain, ![X9, X10, X11, X12, X13]:(((r2_lattice3(X9,X11,X10)|(X10!=k1_yellow_0(X9,X11)|~r1_yellow_0(X9,X11))|~m1_subset_1(X10,u1_struct_0(X9))|(~v5_orders_2(X9)|~l1_orders_2(X9)))&(~m1_subset_1(X12,u1_struct_0(X9))|(~r2_lattice3(X9,X11,X12)|r1_orders_2(X9,X10,X12))|(X10!=k1_yellow_0(X9,X11)|~r1_yellow_0(X9,X11))|~m1_subset_1(X10,u1_struct_0(X9))|(~v5_orders_2(X9)|~l1_orders_2(X9))))&(((X10=k1_yellow_0(X9,X13)|(m1_subset_1(esk3_3(X9,X10,X13),u1_struct_0(X9))|~r2_lattice3(X9,X13,X10))|~m1_subset_1(X10,u1_struct_0(X9))|(~v5_orders_2(X9)|~l1_orders_2(X9)))&(r1_yellow_0(X9,X13)|(m1_subset_1(esk3_3(X9,X10,X13),u1_struct_0(X9))|~r2_lattice3(X9,X13,X10))|~m1_subset_1(X10,u1_struct_0(X9))|(~v5_orders_2(X9)|~l1_orders_2(X9))))&(((X10=k1_yellow_0(X9,X13)|(r2_lattice3(X9,X13,esk3_3(X9,X10,X13))|~r2_lattice3(X9,X13,X10))|~m1_subset_1(X10,u1_struct_0(X9))|(~v5_orders_2(X9)|~l1_orders_2(X9)))&(r1_yellow_0(X9,X13)|(r2_lattice3(X9,X13,esk3_3(X9,X10,X13))|~r2_lattice3(X9,X13,X10))|~m1_subset_1(X10,u1_struct_0(X9))|(~v5_orders_2(X9)|~l1_orders_2(X9))))&((X10=k1_yellow_0(X9,X13)|(~r1_orders_2(X9,X10,esk3_3(X9,X10,X13))|~r2_lattice3(X9,X13,X10))|~m1_subset_1(X10,u1_struct_0(X9))|(~v5_orders_2(X9)|~l1_orders_2(X9)))&(r1_yellow_0(X9,X13)|(~r1_orders_2(X9,X10,esk3_3(X9,X10,X13))|~r2_lattice3(X9,X13,X10))|~m1_subset_1(X10,u1_struct_0(X9))|(~v5_orders_2(X9)|~l1_orders_2(X9))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_yellow_0])])])])])])).
cnf(c_0_16, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, plain, ![X15, X16]:((m1_subset_1(esk4_2(X15,X16),X15)|X15=X16|~m1_subset_1(X16,k1_zfmisc_1(X15)))&(~r2_hidden(esk4_2(X15,X16),X16)|X15=X16|~m1_subset_1(X16,k1_zfmisc_1(X15)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t49_subset_1])])])])).
fof(c_0_19, plain, ![X18, X19, X20, X21]:((~v13_waybel_0(X19,X18)|(~m1_subset_1(X20,u1_struct_0(X18))|(~m1_subset_1(X21,u1_struct_0(X18))|(~r2_hidden(X20,X19)|~r1_orders_2(X18,X20,X21)|r2_hidden(X21,X19))))|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|~l1_orders_2(X18))&((m1_subset_1(esk5_2(X18,X19),u1_struct_0(X18))|v13_waybel_0(X19,X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|~l1_orders_2(X18))&((m1_subset_1(esk6_2(X18,X19),u1_struct_0(X18))|v13_waybel_0(X19,X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|~l1_orders_2(X18))&(((r2_hidden(esk5_2(X18,X19),X19)|v13_waybel_0(X19,X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|~l1_orders_2(X18))&(r1_orders_2(X18,esk5_2(X18,X19),esk6_2(X18,X19))|v13_waybel_0(X19,X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|~l1_orders_2(X18)))&(~r2_hidden(esk6_2(X18,X19),X19)|v13_waybel_0(X19,X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|~l1_orders_2(X18)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d20_waybel_0])])])])])).
fof(c_0_20, plain, ![X24, X25, X26]:(~r2_hidden(X24,X25)|~m1_subset_1(X25,k1_zfmisc_1(X26))|m1_subset_1(X24,X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_21, plain, (r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~r1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, negated_conjecture, (m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_23, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_24, plain, (m1_subset_1(esk4_2(X1,X2),X1)|X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_26, plain, ![X27, X28]:((r2_lattice3(X27,k1_xboole_0,X28)|~m1_subset_1(X28,u1_struct_0(X27))|~l1_orders_2(X27))&(r1_lattice3(X27,k1_xboole_0,X28)|~m1_subset_1(X28,u1_struct_0(X27))|~l1_orders_2(X27))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_yellow_0])])])])).
fof(c_0_27, plain, ![X29]:((r1_yellow_0(X29,k1_xboole_0)|(v2_struct_0(X29)|(~v5_orders_2(X29)|(~v1_yellow_0(X29)|~l1_orders_2(X29)))))&(r2_yellow_0(X29,u1_struct_0(X29))|(v2_struct_0(X29)|(~v5_orders_2(X29)|(~v1_yellow_0(X29)|~l1_orders_2(X29)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t42_yellow_0])])])])).
fof(c_0_28, plain, ![X30]:(~l1_orders_2(X30)|k3_yellow_0(X30)=k1_yellow_0(X30,k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_yellow_0])])).
cnf(c_0_29, plain, (r2_hidden(X4,X1)|~v13_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,X1)|~r1_orders_2(X2,X3,X4)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_30, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_31, negated_conjecture, (r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),X2)|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,X3)|~r2_lattice3(esk1_0,X3,X2)|~r1_yellow_0(esk1_0,X3)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_17]), c_0_23])])).
cnf(c_0_32, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|m1_subset_1(esk4_2(u1_struct_0(esk1_0),esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_33, plain, (r2_lattice3(X1,k1_xboole_0,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (r1_yellow_0(X1,k1_xboole_0)|v2_struct_0(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35, negated_conjecture, (v1_yellow_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_36, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_37, plain, (k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_38, plain, ![X31, X32]:((~v1_subset_1(X32,X31)|X32!=X31|~m1_subset_1(X32,k1_zfmisc_1(X31)))&(X32=X31|v1_subset_1(X32,X31)|~m1_subset_1(X32,k1_zfmisc_1(X31)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_subset_1])])])])).
cnf(c_0_39, plain, (r2_hidden(X1,X2)|~r1_orders_2(X3,X4,X1)|~r2_hidden(X4,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~v13_waybel_0(X2,X3)|~l1_orders_2(X3)), inference(csr,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_40, negated_conjecture, (v13_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_41, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),esk4_2(u1_struct_0(esk1_0),esk2_0))|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,X2)|~r2_lattice3(esk1_0,X2,esk4_2(u1_struct_0(esk1_0),esk2_0))|~r1_yellow_0(esk1_0,X2)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_42, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_lattice3(esk1_0,k1_xboole_0,esk4_2(u1_struct_0(esk1_0),esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_32]), c_0_17])])).
cnf(c_0_43, negated_conjecture, (r1_yellow_0(esk1_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_17]), c_0_35]), c_0_23])]), c_0_36])).
cnf(c_0_44, negated_conjecture, (r2_hidden(k3_yellow_0(esk1_0),esk2_0)|~v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_45, negated_conjecture, (k3_yellow_0(esk1_0)=k1_yellow_0(esk1_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_37, c_0_17])).
cnf(c_0_46, plain, (X1=X2|v1_subset_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_47, negated_conjecture, (r2_hidden(X1,esk2_0)|~r1_orders_2(esk1_0,X2,X1)|~r2_hidden(X2,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_25]), c_0_40]), c_0_17])])).
cnf(c_0_48, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),esk4_2(u1_struct_0(esk1_0),esk2_0))|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43])])).
cnf(c_0_49, negated_conjecture, (r2_hidden(k1_yellow_0(esk1_0,k1_xboole_0),esk2_0)|~v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(rw,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_50, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_46, c_0_25])).
fof(c_0_51, plain, ![X33]:(m1_subset_1(esk7_1(X33),k1_zfmisc_1(X33))&~v1_subset_1(esk7_1(X33),X33)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[rc3_subset_1])])])).
cnf(c_0_52, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_hidden(esk4_2(u1_struct_0(esk1_0),esk2_0),esk2_0)|~r1_orders_2(esk1_0,X1,esk4_2(u1_struct_0(esk1_0),esk2_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_47, c_0_32])).
cnf(c_0_53, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,k1_xboole_0),esk4_2(u1_struct_0(esk1_0),esk2_0))), inference(er,[status(thm)],[c_0_48])).
cnf(c_0_54, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_hidden(k1_yellow_0(esk1_0,k1_xboole_0),esk2_0)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_55, plain, (m1_subset_1(esk7_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_56, plain, (~v1_subset_1(esk7_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
fof(c_0_57, plain, ![X35, X36]:(~m1_subset_1(X35,X36)|(v1_xboole_0(X36)|r2_hidden(X35,X36))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_58, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k3_yellow_0(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_59, plain, (X1=X2|~r2_hidden(esk4_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_60, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_hidden(esk4_2(u1_struct_0(esk1_0),esk2_0),esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54])).
cnf(c_0_61, plain, (esk7_1(X1)=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_55]), c_0_56])).
cnf(c_0_62, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_63, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k1_yellow_0(esk1_0,k1_xboole_0),esk2_0)), inference(rw,[status(thm)],[c_0_58, c_0_45])).
cnf(c_0_64, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_25])])).
cnf(c_0_65, plain, (~v1_subset_1(X1,X1)), inference(rw,[status(thm)],[c_0_56, c_0_61])).
cnf(c_0_66, negated_conjecture, (r2_hidden(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_62, c_0_22])).
cnf(c_0_67, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_68, negated_conjecture, (~r2_hidden(k1_yellow_0(esk1_0,k1_xboole_0),esk2_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_64]), c_0_65])).
cnf(c_0_69, negated_conjecture, (r2_hidden(k1_yellow_0(esk1_0,X1),esk2_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_64]), c_0_64]), c_0_67])).
cnf(c_0_70, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_69])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 71
# Proof object clause steps            : 46
# Proof object formula steps           : 25
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 29
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 40
# Processed clauses                    : 234
# ...of these trivial                  : 0
# ...subsumed                          : 26
# ...remaining for further processing  : 208
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 108
# Generated clauses                    : 232
# ...of the previous two non-trivial   : 230
# Contextual simplify-reflections      : 5
# Paramodulations                      : 228
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 57
#    Positive orientable unit clauses  : 16
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 38
# Current number of unprocessed clauses: 63
# ...number of literals in the above   : 294
# Current number of archived formulas  : 0
# Current number of archived clauses   : 150
# Clause-clause subsumption calls (NU) : 4115
# Rec. Clause-clause subsumption calls : 893
# Non-unit clause-clause subsumptions  : 31
# Unit Clause-clause subsumption calls : 66
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 16
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 10354

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.003 s
# Total time               : 0.019 s
# Maximum resident set size: 28652 pages
