# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', d10_xboole_0)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', d1_tarski)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', t72_enumset1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', l51_zfmisc_1)).
fof(t12_ordinal1, conjecture, ![X1, X2]:(k1_ordinal1(X1)=k1_ordinal1(X2)=>X1=X2), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', t12_ordinal1)).
fof(d1_ordinal1, axiom, ![X1]:k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1)), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', d1_ordinal1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', t12_setfam_1)).
fof(t65_zfmisc_1, axiom, ![X1, X2]:(k4_xboole_0(X1,k1_tarski(X2))=X1<=>~(r2_hidden(X2,X1))), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', t65_zfmisc_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', l32_xboole_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', t7_xboole_1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', t4_xboole_0)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', t48_xboole_1)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', t41_xboole_1)).
fof(idempotence_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', idempotence_k2_xboole_0)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', t7_ordinal1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', t36_xboole_1)).
fof(t10_ordinal1, axiom, ![X1]:r2_hidden(X1,k1_ordinal1(X1)), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', t10_ordinal1)).
fof(t44_zfmisc_1, axiom, ![X1, X2, X3]:~((k1_tarski(X1)=k2_xboole_0(X2,X3)&(~(X2=X3)&(~(X2=k1_xboole_0)&~(X3=k1_xboole_0))))), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', t44_zfmisc_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', t12_xboole_1)).
fof(t43_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3)), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', t43_xboole_1)).
fof(t5_xboole_0, axiom, ![X1, X2, X3]:~((r1_xboole_0(X1,X2)&(r2_hidden(X3,k2_xboole_0(X1,X2))&(~((r2_hidden(X3,X1)&~(r2_hidden(X3,X2))))&~((r2_hidden(X3,X2)&~(r2_hidden(X3,X1)))))))), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', t5_xboole_0)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', symmetry_r1_xboole_0)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', d5_xboole_0)).
fof(antisymmetry_r2_hidden, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>~(r2_hidden(X2,X1))), file('data/MPTP/proofs/09118f033d4f4c7aa83304cd17f5bc31.E_input', antisymmetry_r2_hidden)).
fof(c_0_27, plain, ![X5, X6]:(((r1_tarski(X5,X6)|X5!=X6)&(r1_tarski(X6,X5)|X5!=X6))&(~r1_tarski(X5,X6)|~r1_tarski(X6,X5)|X5=X6)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_28, plain, ![X7, X8, X9, X10, X11, X12]:(((~r2_hidden(X9,X8)|X9=X7|X8!=k1_tarski(X7))&(X10!=X7|r2_hidden(X10,X8)|X8!=k1_tarski(X7)))&((~r2_hidden(esk1_2(X11,X12),X12)|esk1_2(X11,X12)!=X11|X12=k1_tarski(X11))&(r2_hidden(esk1_2(X11,X12),X12)|esk1_2(X11,X12)=X11|X12=k1_tarski(X11)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_29, plain, ![X14]:k2_tarski(X14,X14)=k1_tarski(X14), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_30, plain, ![X15, X16]:k1_enumset1(X15,X15,X16)=k2_tarski(X15,X16), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_31, plain, ![X17, X18, X19]:k2_enumset1(X17,X17,X18,X19)=k1_enumset1(X17,X18,X19), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_32, plain, ![X20, X21, X22, X23]:k3_enumset1(X20,X20,X21,X22,X23)=k2_enumset1(X20,X21,X22,X23), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_33, plain, ![X24, X25]:k3_tarski(k2_tarski(X24,X25))=k2_xboole_0(X24,X25), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_34, negated_conjecture, ~(![X1, X2]:(k1_ordinal1(X1)=k1_ordinal1(X2)=>X1=X2)), inference(assume_negation,[status(cth)],[t12_ordinal1])).
fof(c_0_35, plain, ![X28]:k1_ordinal1(X28)=k2_xboole_0(X28,k1_tarski(X28)), inference(variable_rename,[status(thm)],[d1_ordinal1])).
fof(c_0_36, plain, ![X29, X30]:k1_setfam_1(k2_tarski(X29,X30))=k3_xboole_0(X29,X30), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_37, plain, ![X31, X32]:((k4_xboole_0(X31,k1_tarski(X32))!=X31|~r2_hidden(X32,X31))&(r2_hidden(X32,X31)|k4_xboole_0(X31,k1_tarski(X32))=X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t65_zfmisc_1])])])).
fof(c_0_38, plain, ![X33, X34]:((k4_xboole_0(X33,X34)!=k1_xboole_0|r1_tarski(X33,X34))&(~r1_tarski(X33,X34)|k4_xboole_0(X33,X34)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_39, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_40, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_41, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_42, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_43, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_44, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_45, plain, ![X35, X36]:r1_tarski(X35,k2_xboole_0(X35,X36)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_46, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_47, negated_conjecture, (k1_ordinal1(esk2_0)=k1_ordinal1(esk3_0)&esk2_0!=esk3_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_34])])])).
cnf(c_0_48, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_49, plain, ![X37, X38, X40, X41, X42]:((r1_xboole_0(X37,X38)|r2_hidden(esk4_2(X37,X38),k3_xboole_0(X37,X38)))&(~r2_hidden(X42,k3_xboole_0(X40,X41))|~r1_xboole_0(X40,X41))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
cnf(c_0_50, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_51, plain, ![X43, X44]:k4_xboole_0(X43,k4_xboole_0(X43,X44))=k3_xboole_0(X43,X44), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_52, plain, ![X45, X46, X47]:k4_xboole_0(k4_xboole_0(X45,X46),X47)=k4_xboole_0(X45,k2_xboole_0(X46,X47)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
fof(c_0_53, plain, ![X48]:k2_xboole_0(X48,X48)=X48, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k2_xboole_0])])).
fof(c_0_54, plain, ![X49, X50]:(~r2_hidden(X49,X50)|~r1_tarski(X50,X49)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_55, plain, ![X51, X52]:r1_tarski(k4_xboole_0(X51,X52),X51), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_56, plain, ![X53]:r2_hidden(X53,k1_ordinal1(X53)), inference(variable_rename,[status(thm)],[t10_ordinal1])).
fof(c_0_57, plain, ![X54, X55, X56]:(k1_tarski(X54)!=k2_xboole_0(X55,X56)|(X55=X56|(X55=k1_xboole_0|X56=k1_xboole_0))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t44_zfmisc_1])])])).
fof(c_0_58, plain, ![X57, X58]:(~r1_tarski(X57,X58)|k2_xboole_0(X57,X58)=X58), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_59, plain, (k4_xboole_0(X1,k1_tarski(X2))!=X1|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_60, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_61, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_39])).
cnf(c_0_62, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k3_enumset1(X2,X2,X2,X2,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_41]), c_0_42]), c_0_43]), c_0_44])).
fof(c_0_63, plain, ![X59, X60, X61]:(~r1_tarski(X59,k2_xboole_0(X60,X61))|r1_tarski(k4_xboole_0(X59,X60),X61)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
cnf(c_0_64, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_65, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_46, c_0_42])).
cnf(c_0_66, negated_conjecture, (k1_ordinal1(esk2_0)=k1_ordinal1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_67, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k2_tarski(X1,X1))), inference(rw,[status(thm)],[c_0_48, c_0_41])).
cnf(c_0_68, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk4_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_69, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_50, c_0_42])).
cnf(c_0_70, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_71, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_72, plain, (k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_73, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_74, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
fof(c_0_75, plain, ![X62, X63, X64]:(((r2_hidden(X64,X63)|r2_hidden(X64,X62)|~r2_hidden(X64,k2_xboole_0(X62,X63))|~r1_xboole_0(X62,X63))&(~r2_hidden(X64,X62)|r2_hidden(X64,X62)|~r2_hidden(X64,k2_xboole_0(X62,X63))|~r1_xboole_0(X62,X63)))&((r2_hidden(X64,X63)|~r2_hidden(X64,X63)|~r2_hidden(X64,k2_xboole_0(X62,X63))|~r1_xboole_0(X62,X63))&(~r2_hidden(X64,X62)|~r2_hidden(X64,X63)|~r2_hidden(X64,k2_xboole_0(X62,X63))|~r1_xboole_0(X62,X63)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t5_xboole_0])])])])).
cnf(c_0_76, plain, (r2_hidden(X1,k1_ordinal1(X1))), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_77, plain, (X2=X3|X2=k1_xboole_0|X3=k1_xboole_0|k1_tarski(X1)!=k2_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_78, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_79, plain, (k4_xboole_0(X1,k3_enumset1(X2,X2,X2,X2,X2))!=X1|~r2_hidden(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_41]), c_0_42]), c_0_43]), c_0_44])).
cnf(c_0_80, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_81, plain, (r2_hidden(X1,k3_enumset1(X1,X1,X1,X1,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_62])])).
cnf(c_0_82, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_83, plain, (r1_tarski(X1,k3_tarski(k3_enumset1(X1,X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_65]), c_0_43]), c_0_44])).
cnf(c_0_84, negated_conjecture, (k3_tarski(k3_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,k3_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)))=k3_tarski(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_67]), c_0_67]), c_0_42]), c_0_42]), c_0_65]), c_0_65]), c_0_43]), c_0_43]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_44]), c_0_44])).
cnf(c_0_85, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk4_2(X1,X2),k1_setfam_1(k3_enumset1(X1,X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_69]), c_0_43]), c_0_44])).
cnf(c_0_86, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_setfam_1(k3_enumset1(X1,X1,X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_69]), c_0_43]), c_0_44])).
cnf(c_0_87, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k3_tarski(k3_enumset1(X2,X2,X2,X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_65]), c_0_43]), c_0_44])).
cnf(c_0_88, plain, (k3_tarski(k3_enumset1(X1,X1,X1,X1,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_65]), c_0_43]), c_0_44])).
cnf(c_0_89, plain, (~r2_hidden(X1,k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_90, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k2_xboole_0(X3,X2))|~r1_xboole_0(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_91, plain, (r2_hidden(X1,k3_tarski(k3_enumset1(X1,X1,X1,X1,k3_enumset1(X1,X1,X1,X1,X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_67]), c_0_42]), c_0_65]), c_0_43]), c_0_43]), c_0_44]), c_0_44])).
cnf(c_0_92, plain, (X2=X3|X3=k1_xboole_0|X2=k1_xboole_0|k3_enumset1(X1,X1,X1,X1,X1)!=k3_tarski(k3_enumset1(X2,X2,X2,X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_41]), c_0_42]), c_0_65]), c_0_43]), c_0_43]), c_0_44]), c_0_44])).
cnf(c_0_93, plain, (k3_tarski(k3_enumset1(X1,X1,X1,X1,X2))=X2|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_65]), c_0_43]), c_0_44])).
cnf(c_0_94, plain, (k3_enumset1(X1,X1,X1,X1,X1)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_81])])).
cnf(c_0_95, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k3_tarski(k3_enumset1(X2,X2,X2,X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_65]), c_0_43]), c_0_44])).
cnf(c_0_96, negated_conjecture, (r1_tarski(esk2_0,k3_tarski(k3_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,k3_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))))), inference(spm,[status(thm)],[c_0_83, c_0_84])).
fof(c_0_97, plain, ![X65, X66]:(~r1_xboole_0(X65,X66)|r1_xboole_0(X66,X65)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
cnf(c_0_98, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk4_2(X1,X2),k4_xboole_0(X1,k4_xboole_0(X1,X2)))), inference(rw,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_99, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_100, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_89, c_0_80])).
cnf(c_0_101, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_102, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_103, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X2)|~r1_xboole_0(X3,X2)|~r2_hidden(X1,k3_tarski(k3_enumset1(X3,X3,X3,X3,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_65]), c_0_43]), c_0_44])).
cnf(c_0_104, negated_conjecture, (r2_hidden(esk2_0,k3_tarski(k3_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,k3_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))))), inference(spm,[status(thm)],[c_0_91, c_0_84])).
cnf(c_0_105, plain, (X1=k3_enumset1(X2,X2,X2,X2,X2)|X1=k1_xboole_0|~r1_tarski(X1,k3_enumset1(X2,X2,X2,X2,X2))), inference(sr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93])]), c_0_94])).
cnf(c_0_106, negated_conjecture, (r1_tarski(k4_xboole_0(esk2_0,esk3_0),k3_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))), inference(spm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_107, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_108, plain, (r1_xboole_0(k4_xboole_0(X1,X2),X2)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_80]), c_0_100])).
cnf(c_0_109, negated_conjecture, (r1_tarski(k4_xboole_0(X1,esk2_0),k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))|~r1_tarski(X1,k3_tarski(k3_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,k3_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))))), inference(spm,[status(thm)],[c_0_95, c_0_84])).
cnf(c_0_110, plain, (X1=X2|k4_xboole_0(X2,X1)!=k1_xboole_0|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_111, negated_conjecture, (r2_hidden(esk2_0,k3_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))|r2_hidden(esk2_0,esk3_0)|~r1_xboole_0(esk3_0,k3_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))), inference(spm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_112, negated_conjecture, (k3_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)=k4_xboole_0(esk2_0,esk3_0)|k4_xboole_0(esk2_0,esk3_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_113, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_114, negated_conjecture, (r1_tarski(k4_xboole_0(esk3_0,esk2_0),k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))), inference(spm,[status(thm)],[c_0_109, c_0_83])).
cnf(c_0_115, plain, (X1=X2|k4_xboole_0(X2,X1)!=k1_xboole_0|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_110, c_0_102])).
cnf(c_0_116, negated_conjecture, (k4_xboole_0(esk2_0,esk3_0)=k1_xboole_0|r2_hidden(esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_113])]), c_0_89])).
cnf(c_0_117, negated_conjecture, (esk2_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_118, plain, (k4_xboole_0(X1,X2)!=k1_xboole_0|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_73, c_0_102])).
fof(c_0_119, plain, ![X67, X68, X69, X70, X71, X72, X73, X74]:((((r2_hidden(X70,X67)|~r2_hidden(X70,X69)|X69!=k4_xboole_0(X67,X68))&(~r2_hidden(X70,X68)|~r2_hidden(X70,X69)|X69!=k4_xboole_0(X67,X68)))&(~r2_hidden(X71,X67)|r2_hidden(X71,X68)|r2_hidden(X71,X69)|X69!=k4_xboole_0(X67,X68)))&((~r2_hidden(esk5_3(X72,X73,X74),X74)|(~r2_hidden(esk5_3(X72,X73,X74),X72)|r2_hidden(esk5_3(X72,X73,X74),X73))|X74=k4_xboole_0(X72,X73))&((r2_hidden(esk5_3(X72,X73,X74),X72)|r2_hidden(esk5_3(X72,X73,X74),X74)|X74=k4_xboole_0(X72,X73))&(~r2_hidden(esk5_3(X72,X73,X74),X73)|r2_hidden(esk5_3(X72,X73,X74),X74)|X74=k4_xboole_0(X72,X73))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_120, negated_conjecture, (k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)=k4_xboole_0(esk3_0,esk2_0)|k4_xboole_0(esk3_0,esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_105, c_0_114])).
cnf(c_0_121, negated_conjecture, (k4_xboole_0(esk3_0,esk2_0)!=k1_xboole_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_117]), c_0_118])).
cnf(c_0_122, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_123, negated_conjecture, (k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)=k4_xboole_0(esk3_0,esk2_0)), inference(sr,[status(thm)],[c_0_120, c_0_121])).
fof(c_0_124, plain, ![X76, X77]:(~r2_hidden(X76,X77)|~r2_hidden(X77,X76)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[antisymmetry_r2_hidden])])])).
cnf(c_0_125, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_122])).
cnf(c_0_126, negated_conjecture, (r2_hidden(esk2_0,k4_xboole_0(esk3_0,esk2_0))), inference(spm,[status(thm)],[c_0_81, c_0_123])).
cnf(c_0_127, negated_conjecture, (k3_tarski(k3_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,k3_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)))=k3_tarski(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,k4_xboole_0(esk3_0,esk2_0)))), inference(rw,[status(thm)],[c_0_84, c_0_123])).
cnf(c_0_128, plain, (~r2_hidden(X1,X2)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_129, negated_conjecture, (r2_hidden(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_125, c_0_126])).
cnf(c_0_130, negated_conjecture, (r2_hidden(esk3_0,k3_tarski(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,k4_xboole_0(esk3_0,esk2_0))))), inference(spm,[status(thm)],[c_0_91, c_0_127])).
cnf(c_0_131, negated_conjecture, (~r2_hidden(esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_128, c_0_129])).
cnf(c_0_132, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_130]), c_0_113])]), c_0_131]), c_0_89]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 133
# Proof object clause steps            : 78
# Proof object formula steps           : 55
# Proof object conjectures             : 23
# Proof object clause conjectures      : 20
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 30
# Proof object initial formulas used   : 27
# Proof object generating inferences   : 26
# Proof object simplifying inferences  : 85
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 27
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 44
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 35
# Processed clauses                    : 1039
# ...of these trivial                  : 15
# ...subsumed                          : 695
# ...remaining for further processing  : 329
# Other redundant clauses eliminated   : 31
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 20
# Backward-rewritten                   : 91
# Generated clauses                    : 3327
# ...of the previous two non-trivial   : 2850
# Contextual simplify-reflections      : 1
# Paramodulations                      : 3287
# Factorizations                       : 2
# Equation resolutions                 : 31
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 169
#    Positive orientable unit clauses  : 39
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 31
#    Non-unit-clauses                  : 99
# Current number of unprocessed clauses: 1817
# ...number of literals in the above   : 5135
# Current number of archived formulas  : 0
# Current number of archived clauses   : 160
# Clause-clause subsumption calls (NU) : 2488
# Rec. Clause-clause subsumption calls : 2011
# Non-unit clause-clause subsumptions  : 155
# Unit Clause-clause subsumption calls : 501
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 45
# BW rewrite match successes           : 13
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 40747

# -------------------------------------------------
# User time                : 0.030 s
# System time              : 0.010 s
# Total time               : 0.040 s
# Maximum resident set size: 28636 pages
