# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/0931f8ae496841d28a20afa6423ba817.E_input', d4_relat_1)).
fof(d7_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(X2=k4_relat_1(X1)<=>![X3, X4]:(r2_hidden(k4_tarski(X3,X4),X2)<=>r2_hidden(k4_tarski(X4,X3),X1))))), file('data/MPTP/proofs/0931f8ae496841d28a20afa6423ba817.E_input', d7_relat_1)).
fof(dt_k4_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>v1_relat_1(k4_relat_1(X1))), file('data/MPTP/proofs/0931f8ae496841d28a20afa6423ba817.E_input', dt_k4_relat_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/0931f8ae496841d28a20afa6423ba817.E_input', d5_xboole_0)).
fof(t4_boole, axiom, ![X1]:k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0, file('data/MPTP/proofs/0931f8ae496841d28a20afa6423ba817.E_input', t4_boole)).
fof(t56_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>(![X2, X3]:~(r2_hidden(k4_tarski(X2,X3),X1))=>X1=k1_xboole_0)), file('data/MPTP/proofs/0931f8ae496841d28a20afa6423ba817.E_input', t56_relat_1)).
fof(t7_tarski, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&r2_hidden(X4,X3))))))), file('data/MPTP/proofs/0931f8ae496841d28a20afa6423ba817.E_input', t7_tarski)).
fof(involutiveness_k4_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k4_relat_1(k4_relat_1(X1))=X1), file('data/MPTP/proofs/0931f8ae496841d28a20afa6423ba817.E_input', involutiveness_k4_relat_1)).
fof(t66_relat_1, conjecture, k4_relat_1(k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/0931f8ae496841d28a20afa6423ba817.E_input', t66_relat_1)).
fof(t60_relat_1, axiom, (k1_relat_1(k1_xboole_0)=k1_xboole_0&k2_relat_1(k1_xboole_0)=k1_xboole_0), file('data/MPTP/proofs/0931f8ae496841d28a20afa6423ba817.E_input', t60_relat_1)).
fof(cc1_relat_1, axiom, ![X1]:(v1_xboole_0(X1)=>v1_relat_1(X1)), file('data/MPTP/proofs/0931f8ae496841d28a20afa6423ba817.E_input', cc1_relat_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/0931f8ae496841d28a20afa6423ba817.E_input', fc1_xboole_0)).
fof(c_0_12, plain, ![X10, X11, X12, X14, X15, X16, X17, X19]:(((~r2_hidden(X12,X11)|r2_hidden(k4_tarski(X12,esk3_3(X10,X11,X12)),X10)|X11!=k1_relat_1(X10))&(~r2_hidden(k4_tarski(X14,X15),X10)|r2_hidden(X14,X11)|X11!=k1_relat_1(X10)))&((~r2_hidden(esk4_2(X16,X17),X17)|~r2_hidden(k4_tarski(esk4_2(X16,X17),X19),X16)|X17=k1_relat_1(X16))&(r2_hidden(esk4_2(X16,X17),X17)|r2_hidden(k4_tarski(esk4_2(X16,X17),esk5_2(X16,X17)),X16)|X17=k1_relat_1(X16)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
fof(c_0_13, plain, ![X32, X33, X34, X35, X36, X37]:(((~r2_hidden(k4_tarski(X34,X35),X33)|r2_hidden(k4_tarski(X35,X34),X32)|X33!=k4_relat_1(X32)|~v1_relat_1(X33)|~v1_relat_1(X32))&(~r2_hidden(k4_tarski(X37,X36),X32)|r2_hidden(k4_tarski(X36,X37),X33)|X33!=k4_relat_1(X32)|~v1_relat_1(X33)|~v1_relat_1(X32)))&((~r2_hidden(k4_tarski(esk7_2(X32,X33),esk8_2(X32,X33)),X33)|~r2_hidden(k4_tarski(esk8_2(X32,X33),esk7_2(X32,X33)),X32)|X33=k4_relat_1(X32)|~v1_relat_1(X33)|~v1_relat_1(X32))&(r2_hidden(k4_tarski(esk7_2(X32,X33),esk8_2(X32,X33)),X33)|r2_hidden(k4_tarski(esk8_2(X32,X33),esk7_2(X32,X33)),X32)|X33=k4_relat_1(X32)|~v1_relat_1(X33)|~v1_relat_1(X32)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_relat_1])])])])])])).
fof(c_0_14, plain, ![X5]:(~v1_relat_1(X5)|v1_relat_1(k4_relat_1(X5))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_relat_1])])).
fof(c_0_15, plain, ![X21, X22, X23, X24, X25, X26, X27, X28]:((((r2_hidden(X24,X21)|~r2_hidden(X24,X23)|X23!=k4_xboole_0(X21,X22))&(~r2_hidden(X24,X22)|~r2_hidden(X24,X23)|X23!=k4_xboole_0(X21,X22)))&(~r2_hidden(X25,X21)|r2_hidden(X25,X22)|r2_hidden(X25,X23)|X23!=k4_xboole_0(X21,X22)))&((~r2_hidden(esk6_3(X26,X27,X28),X28)|(~r2_hidden(esk6_3(X26,X27,X28),X26)|r2_hidden(esk6_3(X26,X27,X28),X27))|X28=k4_xboole_0(X26,X27))&((r2_hidden(esk6_3(X26,X27,X28),X26)|r2_hidden(esk6_3(X26,X27,X28),X28)|X28=k4_xboole_0(X26,X27))&(~r2_hidden(esk6_3(X26,X27,X28),X27)|r2_hidden(esk6_3(X26,X27,X28),X28)|X28=k4_xboole_0(X26,X27))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_16, plain, (r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, plain, (r2_hidden(k4_tarski(X2,X1),X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k4_relat_1(X3)|~v1_relat_1(X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, plain, (v1_relat_1(k4_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_20, plain, ![X30]:k4_xboole_0(k1_xboole_0,X30)=k1_xboole_0, inference(variable_rename,[status(thm)],[t4_boole])).
cnf(c_0_21, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (r2_hidden(k4_tarski(X1,X2),k4_relat_1(X3))|~r2_hidden(k4_tarski(X2,X1),X3)|~v1_relat_1(X3)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_17]), c_0_18])).
fof(c_0_23, plain, ![X7]:(~v1_relat_1(X7)|(r2_hidden(k4_tarski(esk1_1(X7),esk2_1(X7)),X7)|X7=k1_xboole_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t56_relat_1])])])])).
cnf(c_0_24, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_26, plain, ![X40, X41, X43]:((r2_hidden(esk9_2(X40,X41),X41)|~r2_hidden(X40,X41))&(~r2_hidden(X43,X41)|~r2_hidden(X43,esk9_2(X40,X41))|~r2_hidden(X40,X41))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tarski])])])])])).
cnf(c_0_27, plain, (r2_hidden(X1,k1_relat_1(k4_relat_1(X2)))|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28, plain, (r2_hidden(k4_tarski(esk1_1(X1),esk2_1(X1)),X1)|X1=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_29, plain, ![X6]:(~v1_relat_1(X6)|k4_relat_1(k4_relat_1(X6))=X6), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k4_relat_1])])).
fof(c_0_30, negated_conjecture, ~(k4_relat_1(k1_xboole_0)=k1_xboole_0), inference(assume_negation,[status(cth)],[t66_relat_1])).
cnf(c_0_31, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_32, plain, (r2_hidden(esk9_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, plain, (X1=k1_xboole_0|r2_hidden(esk2_1(X1),k1_relat_1(k4_relat_1(X1)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34, plain, (k4_relat_1(k4_relat_1(X1))=X1|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_35, negated_conjecture, k4_relat_1(k1_xboole_0)!=k1_xboole_0, inference(fof_simplification,[status(thm)],[c_0_30])).
cnf(c_0_36, plain, (~r2_hidden(esk9_2(X1,k1_xboole_0),X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_37, plain, (k4_relat_1(X1)=k1_xboole_0|r2_hidden(esk2_1(k4_relat_1(X1)),k1_relat_1(X1))|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_18])).
cnf(c_0_38, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
cnf(c_0_39, negated_conjecture, (k4_relat_1(k1_xboole_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_40, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_36, c_0_32])).
fof(c_0_41, plain, ![X31]:(~v1_xboole_0(X31)|v1_relat_1(X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relat_1])])).
cnf(c_0_42, plain, (~v1_relat_1(k1_xboole_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39]), c_0_40])).
cnf(c_0_43, plain, (v1_relat_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_44, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_45, plain, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 46
# Proof object clause steps            : 23
# Proof object formula steps           : 23
# Proof object conjectures             : 4
# Proof object clause conjectures      : 1
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 9
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 25
# Processed clauses                    : 80
# ...of these trivial                  : 0
# ...subsumed                          : 19
# ...remaining for further processing  : 61
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 152
# ...of the previous two non-trivial   : 123
# Contextual simplify-reflections      : 3
# Paramodulations                      : 143
# Factorizations                       : 2
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 53
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 45
# Current number of unprocessed clauses: 68
# ...number of literals in the above   : 226
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1
# Clause-clause subsumption calls (NU) : 399
# Rec. Clause-clause subsumption calls : 262
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 10
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3756

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.004 s
# Total time               : 0.011 s
# Maximum resident set size: 28556 pages
