# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t88_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))=>(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))=X4=>![X6]:(m1_subset_1(X6,u1_struct_0(X4))=>![X7]:(m1_subset_1(X7,u1_struct_0(X3))=>((X6=X7&r1_tmap_1(X3,X2,k3_tmap_1(X1,X2,X4,X3,X5),X7))=>r1_tmap_1(X4,X2,X5,X6)))))))))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', t88_tmap_1)).
fof(t11_tmap_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>(v1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))&m1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)),X1)))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', t11_tmap_1)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', dt_m1_pre_topc)).
fof(free_g1_pre_topc, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>![X3, X4]:(g1_pre_topc(X1,X2)=g1_pre_topc(X3,X4)=>(X1=X3&X2=X4))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', free_g1_pre_topc)).
fof(abstractness_v1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v1_pre_topc(X1)=>X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', abstractness_v1_pre_topc)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', d10_xboole_0)).
fof(t59_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))=>m1_pre_topc(X2,X1))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', t59_pre_topc)).
fof(cc1_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>v2_pre_topc(X2))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', cc1_pre_topc)).
fof(t4_tsep_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X1)=>(r1_tarski(u1_struct_0(X2),u1_struct_0(X3))<=>m1_pre_topc(X2,X3))))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', t4_tsep_1)).
fof(dt_u1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', dt_u1_pre_topc)).
fof(t2_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>m1_pre_topc(X1,X1)), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', t2_tsep_1)).
fof(d1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1))))))))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', d1_pre_topc)).
fof(t65_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>(m1_pre_topc(X1,X2)<=>m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', t65_pre_topc)).
fof(t1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', t1_tsep_1)).
fof(t7_tsep_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X2)=>m1_pre_topc(X3,X1)))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', t7_tsep_1)).
fof(d5_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:(m1_pre_topc(X3,X1)=>![X4]:(m1_pre_topc(X4,X1)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))=>(m1_pre_topc(X4,X3)=>k3_tmap_1(X1,X2,X3,X4,X5)=k2_partfun1(u1_struct_0(X3),u1_struct_0(X2),X5,u1_struct_0(X4)))))))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', d5_tmap_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', t3_subset)).
fof(d5_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>r2_hidden(X2,u1_pre_topc(X1))))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', d5_pre_topc)).
fof(d4_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:(m1_pre_topc(X4,X1)=>k2_tmap_1(X1,X2,X3,X4)=k2_partfun1(u1_struct_0(X1),u1_struct_0(X2),X3,u1_struct_0(X4)))))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', d4_tmap_1)).
fof(t16_tsep_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>((v1_tsep_1(X2,X1)&m1_pre_topc(X2,X1))<=>v3_pre_topc(X3,X1)))))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', t16_tsep_1)).
fof(t67_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>![X4]:((~(v2_struct_0(X4))&(v1_tsep_1(X4,X2)&m1_pre_topc(X4,X2)))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>![X6]:(m1_subset_1(X6,u1_struct_0(X4))=>(X5=X6=>(r1_tmap_1(X2,X1,X3,X5)<=>r1_tmap_1(X4,X1,k2_tmap_1(X2,X1,X3,X4),X6))))))))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', t67_tmap_1)).
fof(t55_pre_topc, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>m1_subset_1(X3,u1_struct_0(X1))))), file('data/MPTP/proofs/096ea797e8db4b3bac5866d41bdcc09e.E_input', t55_pre_topc)).
fof(c_0_22, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))=>(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))=X4=>![X6]:(m1_subset_1(X6,u1_struct_0(X4))=>![X7]:(m1_subset_1(X7,u1_struct_0(X3))=>((X6=X7&r1_tmap_1(X3,X2,k3_tmap_1(X1,X2,X4,X3,X5),X7))=>r1_tmap_1(X4,X2,X5,X6))))))))))), inference(assume_negation,[status(cth)],[t88_tmap_1])).
fof(c_0_23, plain, ![X19, X20]:((v1_pre_topc(g1_pre_topc(u1_struct_0(X20),u1_pre_topc(X20)))|~m1_pre_topc(X20,X19)|~l1_pre_topc(X19))&(m1_pre_topc(g1_pre_topc(u1_struct_0(X20),u1_pre_topc(X20)),X19)|~m1_pre_topc(X20,X19)|~l1_pre_topc(X19))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_tmap_1])])])])).
fof(c_0_24, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk1_0))&((~v2_struct_0(esk4_0)&m1_pre_topc(esk4_0,esk1_0))&((v1_funct_1(esk5_0)&(v1_funct_2(esk5_0,u1_struct_0(esk4_0),u1_struct_0(esk2_0))&m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))))&(g1_pre_topc(u1_struct_0(esk3_0),u1_pre_topc(esk3_0))=esk4_0&(m1_subset_1(esk6_0,u1_struct_0(esk4_0))&(m1_subset_1(esk7_0,u1_struct_0(esk3_0))&((esk6_0=esk7_0&r1_tmap_1(esk3_0,esk2_0,k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk5_0),esk7_0))&~r1_tmap_1(esk4_0,esk2_0,esk5_0,esk6_0)))))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_22])])])])).
fof(c_0_25, plain, ![X36, X37]:(~l1_pre_topc(X36)|(~m1_pre_topc(X37,X36)|l1_pre_topc(X37))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
fof(c_0_26, plain, ![X38, X39, X40, X41]:((X38=X40|g1_pre_topc(X38,X39)!=g1_pre_topc(X40,X41)|~m1_subset_1(X39,k1_zfmisc_1(k1_zfmisc_1(X38))))&(X39=X41|g1_pre_topc(X38,X39)!=g1_pre_topc(X40,X41)|~m1_subset_1(X39,k1_zfmisc_1(k1_zfmisc_1(X38))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_pre_topc])])])])).
fof(c_0_27, plain, ![X15]:(~l1_pre_topc(X15)|(~v1_pre_topc(X15)|X15=g1_pre_topc(u1_struct_0(X15),u1_pre_topc(X15)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_pre_topc])])).
cnf(c_0_28, plain, (v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, negated_conjecture, (g1_pre_topc(u1_struct_0(esk3_0),u1_pre_topc(esk3_0))=esk4_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, negated_conjecture, (m1_pre_topc(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_34, plain, ![X56, X57]:(((r1_tarski(X56,X57)|X56!=X57)&(r1_tarski(X57,X56)|X56!=X57))&(~r1_tarski(X56,X57)|~r1_tarski(X57,X56)|X56=X57)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_35, plain, (X1=X2|g1_pre_topc(X3,X1)!=g1_pre_topc(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X3)))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X1)|~v1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, negated_conjecture, (v1_pre_topc(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30]), c_0_31])])).
cnf(c_0_38, negated_conjecture, (l1_pre_topc(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_31])])).
fof(c_0_39, plain, ![X66, X67]:(~l1_pre_topc(X66)|(~m1_pre_topc(X67,g1_pre_topc(u1_struct_0(X66),u1_pre_topc(X66)))|m1_pre_topc(X67,X66))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t59_pre_topc])])])).
fof(c_0_40, plain, ![X17, X18]:(~v2_pre_topc(X17)|~l1_pre_topc(X17)|(~m1_pre_topc(X18,X17)|v2_pre_topc(X18))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_pre_topc])])])).
fof(c_0_41, plain, ![X53, X54, X55]:((~r1_tarski(u1_struct_0(X54),u1_struct_0(X55))|m1_pre_topc(X54,X55)|~m1_pre_topc(X55,X53)|~m1_pre_topc(X54,X53)|(~v2_pre_topc(X53)|~l1_pre_topc(X53)))&(~m1_pre_topc(X54,X55)|r1_tarski(u1_struct_0(X54),u1_struct_0(X55))|~m1_pre_topc(X55,X53)|~m1_pre_topc(X54,X53)|(~v2_pre_topc(X53)|~l1_pre_topc(X53)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_tsep_1])])])])).
cnf(c_0_42, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43, negated_conjecture, (u1_pre_topc(esk3_0)=X1|esk4_0!=g1_pre_topc(X2,X1)|~m1_subset_1(u1_pre_topc(esk3_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk3_0))))), inference(spm,[status(thm)],[c_0_35, c_0_30])).
cnf(c_0_44, negated_conjecture, (g1_pre_topc(u1_struct_0(esk4_0),u1_pre_topc(esk4_0))=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])])).
fof(c_0_45, plain, ![X16]:(~l1_pre_topc(X16)|m1_subset_1(u1_pre_topc(X16),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X16))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_pre_topc])])).
cnf(c_0_46, plain, (m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_47, negated_conjecture, (l1_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_29]), c_0_31])])).
fof(c_0_48, plain, ![X60]:(~l1_pre_topc(X60)|m1_pre_topc(X60,X60)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tsep_1])])).
fof(c_0_49, plain, ![X29, X30, X31, X32]:(((r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|~v2_pre_topc(X29)|~l1_pre_topc(X29))&((~m1_subset_1(X30,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X29))))|(~r1_tarski(X30,u1_pre_topc(X29))|r2_hidden(k5_setfam_1(u1_struct_0(X29),X30),u1_pre_topc(X29)))|~v2_pre_topc(X29)|~l1_pre_topc(X29))&(~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X29)))|(~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X29)))|(~r2_hidden(X31,u1_pre_topc(X29))|~r2_hidden(X32,u1_pre_topc(X29))|r2_hidden(k9_subset_1(u1_struct_0(X29),X31,X32),u1_pre_topc(X29))))|~v2_pre_topc(X29)|~l1_pre_topc(X29))))&(((m1_subset_1(esk9_1(X29),k1_zfmisc_1(u1_struct_0(X29)))|m1_subset_1(esk8_1(X29),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X29))))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29))&((m1_subset_1(esk10_1(X29),k1_zfmisc_1(u1_struct_0(X29)))|m1_subset_1(esk8_1(X29),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X29))))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29))&(((r2_hidden(esk9_1(X29),u1_pre_topc(X29))|m1_subset_1(esk8_1(X29),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X29))))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29))&(r2_hidden(esk10_1(X29),u1_pre_topc(X29))|m1_subset_1(esk8_1(X29),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X29))))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29)))&(~r2_hidden(k9_subset_1(u1_struct_0(X29),esk9_1(X29),esk10_1(X29)),u1_pre_topc(X29))|m1_subset_1(esk8_1(X29),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X29))))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29)))))&(((m1_subset_1(esk9_1(X29),k1_zfmisc_1(u1_struct_0(X29)))|r1_tarski(esk8_1(X29),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29))&((m1_subset_1(esk10_1(X29),k1_zfmisc_1(u1_struct_0(X29)))|r1_tarski(esk8_1(X29),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29))&(((r2_hidden(esk9_1(X29),u1_pre_topc(X29))|r1_tarski(esk8_1(X29),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29))&(r2_hidden(esk10_1(X29),u1_pre_topc(X29))|r1_tarski(esk8_1(X29),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29)))&(~r2_hidden(k9_subset_1(u1_struct_0(X29),esk9_1(X29),esk10_1(X29)),u1_pre_topc(X29))|r1_tarski(esk8_1(X29),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29)))))&((m1_subset_1(esk9_1(X29),k1_zfmisc_1(u1_struct_0(X29)))|~r2_hidden(k5_setfam_1(u1_struct_0(X29),esk8_1(X29)),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29))&((m1_subset_1(esk10_1(X29),k1_zfmisc_1(u1_struct_0(X29)))|~r2_hidden(k5_setfam_1(u1_struct_0(X29),esk8_1(X29)),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29))&(((r2_hidden(esk9_1(X29),u1_pre_topc(X29))|~r2_hidden(k5_setfam_1(u1_struct_0(X29),esk8_1(X29)),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29))&(r2_hidden(esk10_1(X29),u1_pre_topc(X29))|~r2_hidden(k5_setfam_1(u1_struct_0(X29),esk8_1(X29)),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29)))&(~r2_hidden(k9_subset_1(u1_struct_0(X29),esk9_1(X29),esk10_1(X29)),u1_pre_topc(X29))|~r2_hidden(k5_setfam_1(u1_struct_0(X29),esk8_1(X29)),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])).
cnf(c_0_50, plain, (v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_51, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_52, plain, ![X64, X65]:((~m1_pre_topc(X64,X65)|m1_pre_topc(X64,g1_pre_topc(u1_struct_0(X65),u1_pre_topc(X65)))|~l1_pre_topc(X65)|~l1_pre_topc(X64))&(~m1_pre_topc(X64,g1_pre_topc(u1_struct_0(X65),u1_pre_topc(X65)))|m1_pre_topc(X64,X65)|~l1_pre_topc(X65)|~l1_pre_topc(X64))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t65_pre_topc])])])])).
cnf(c_0_53, plain, (m1_pre_topc(X1,X2)|~r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_54, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_55, negated_conjecture, (u1_pre_topc(esk3_0)=u1_pre_topc(esk4_0)|~m1_subset_1(u1_pre_topc(esk3_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk3_0))))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_56, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_57, plain, ![X51, X52]:(~l1_pre_topc(X51)|(~m1_pre_topc(X52,X51)|m1_subset_1(u1_struct_0(X52),k1_zfmisc_1(u1_struct_0(X51))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])).
cnf(c_0_58, negated_conjecture, (m1_pre_topc(X1,esk3_0)|~m1_pre_topc(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_30]), c_0_47])])).
cnf(c_0_59, plain, (m1_pre_topc(X1,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_60, plain, ![X21, X22, X23]:(~v2_pre_topc(X21)|~l1_pre_topc(X21)|(~m1_pre_topc(X22,X21)|(~m1_pre_topc(X23,X22)|m1_pre_topc(X23,X21)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tsep_1])])])).
cnf(c_0_61, plain, (r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_62, negated_conjecture, (v2_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_29]), c_0_31]), c_0_51])])).
cnf(c_0_63, plain, (m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_64, plain, (m1_pre_topc(X1,X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_65, negated_conjecture, (u1_pre_topc(esk3_0)=u1_pre_topc(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_47])])).
fof(c_0_66, plain, ![X24, X25, X26, X27, X28]:(v2_struct_0(X24)|(~v2_pre_topc(X24)|~l1_pre_topc(X24))|(v2_struct_0(X25)|(~v2_pre_topc(X25)|~l1_pre_topc(X25))|(~m1_pre_topc(X26,X24)|(~m1_pre_topc(X27,X24)|(~v1_funct_1(X28)|(~v1_funct_2(X28,u1_struct_0(X26),u1_struct_0(X25))|~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X26),u1_struct_0(X25)))))|(~m1_pre_topc(X27,X26)|k3_tmap_1(X24,X25,X26,X27,X28)=k2_partfun1(u1_struct_0(X26),u1_struct_0(X25),X28,u1_struct_0(X27)))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_tmap_1])])])])).
fof(c_0_67, plain, ![X58, X59]:((~m1_subset_1(X58,k1_zfmisc_1(X59))|r1_tarski(X58,X59))&(~r1_tarski(X58,X59)|m1_subset_1(X58,k1_zfmisc_1(X59)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_68, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_69, negated_conjecture, (m1_pre_topc(esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_38])])).
cnf(c_0_70, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X1,X2)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_71, plain, (m1_pre_topc(X3,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
fof(c_0_72, plain, ![X46, X47]:((~v3_pre_topc(X47,X46)|r2_hidden(X47,u1_pre_topc(X46))|~m1_subset_1(X47,k1_zfmisc_1(u1_struct_0(X46)))|~l1_pre_topc(X46))&(~r2_hidden(X47,u1_pre_topc(X46))|v3_pre_topc(X47,X46)|~m1_subset_1(X47,k1_zfmisc_1(u1_struct_0(X46)))|~l1_pre_topc(X46))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_pre_topc])])])])).
cnf(c_0_73, negated_conjecture, (r2_hidden(u1_struct_0(esk3_0),u1_pre_topc(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_47])])).
cnf(c_0_74, plain, (m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[c_0_63, c_0_32])).
cnf(c_0_75, negated_conjecture, (m1_pre_topc(esk3_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_29]), c_0_31]), c_0_51])])).
cnf(c_0_76, negated_conjecture, (g1_pre_topc(u1_struct_0(esk3_0),u1_pre_topc(esk4_0))=esk4_0), inference(rw,[status(thm)],[c_0_30, c_0_65])).
cnf(c_0_77, plain, (v2_struct_0(X1)|v2_struct_0(X2)|k3_tmap_1(X1,X2,X3,X4,X5)=k2_partfun1(u1_struct_0(X3),u1_struct_0(X2),X5,u1_struct_0(X4))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_pre_topc(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_78, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_79, negated_conjecture, (m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_47])])).
cnf(c_0_80, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X3)|~v2_pre_topc(X3)), inference(csr,[status(thm)],[c_0_70, c_0_71])).
fof(c_0_81, plain, ![X42, X43, X44, X45]:(v2_struct_0(X42)|(~v2_pre_topc(X42)|~l1_pre_topc(X42))|(v2_struct_0(X43)|(~v2_pre_topc(X43)|~l1_pre_topc(X43))|(~v1_funct_1(X44)|(~v1_funct_2(X44,u1_struct_0(X42),u1_struct_0(X43))|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X42),u1_struct_0(X43)))))|(~m1_pre_topc(X45,X42)|k2_tmap_1(X42,X43,X44,X45)=k2_partfun1(u1_struct_0(X42),u1_struct_0(X43),X44,u1_struct_0(X45)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_tmap_1])])])])).
fof(c_0_82, plain, ![X48, X49, X50]:((~v1_tsep_1(X49,X48)|~m1_pre_topc(X49,X48)|v3_pre_topc(X50,X48)|X50!=u1_struct_0(X49)|~m1_subset_1(X50,k1_zfmisc_1(u1_struct_0(X48)))|~m1_pre_topc(X49,X48)|(~v2_pre_topc(X48)|~l1_pre_topc(X48)))&((v1_tsep_1(X49,X48)|~v3_pre_topc(X50,X48)|X50!=u1_struct_0(X49)|~m1_subset_1(X50,k1_zfmisc_1(u1_struct_0(X48)))|~m1_pre_topc(X49,X48)|(~v2_pre_topc(X48)|~l1_pre_topc(X48)))&(m1_pre_topc(X49,X48)|~v3_pre_topc(X50,X48)|X50!=u1_struct_0(X49)|~m1_subset_1(X50,k1_zfmisc_1(u1_struct_0(X48)))|~m1_pre_topc(X49,X48)|(~v2_pre_topc(X48)|~l1_pre_topc(X48))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_tsep_1])])])])).
cnf(c_0_83, plain, (v3_pre_topc(X1,X2)|~r2_hidden(X1,u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_84, negated_conjecture, (r2_hidden(u1_struct_0(esk3_0),u1_pre_topc(esk4_0))), inference(rw,[status(thm)],[c_0_73, c_0_65])).
cnf(c_0_85, negated_conjecture, (m1_pre_topc(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_65]), c_0_76]), c_0_47])])).
cnf(c_0_86, plain, (k3_tmap_1(X1,X2,X3,X4,X5)=k2_partfun1(u1_struct_0(X3),u1_struct_0(X2),X5,u1_struct_0(X4))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_pre_topc(X4,X3)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_77, c_0_71])).
cnf(c_0_87, negated_conjecture, (v1_funct_2(esk5_0,u1_struct_0(esk4_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_88, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_89, negated_conjecture, (v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_90, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_91, negated_conjecture, (v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_92, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_93, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_94, negated_conjecture, (r1_tarski(u1_struct_0(esk4_0),u1_struct_0(esk3_0))), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_95, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk4_0))|~m1_pre_topc(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_33]), c_0_31]), c_0_51])])).
cnf(c_0_96, plain, (v2_struct_0(X1)|v2_struct_0(X2)|k2_tmap_1(X1,X2,X3,X4)=k2_partfun1(u1_struct_0(X1),u1_struct_0(X2),X3,u1_struct_0(X4))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_pre_topc(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_97, negated_conjecture, (v2_pre_topc(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_33]), c_0_31]), c_0_51])])).
cnf(c_0_98, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_99, plain, ![X68, X69, X70, X71, X72, X73]:((~r1_tmap_1(X69,X68,X70,X72)|r1_tmap_1(X71,X68,k2_tmap_1(X69,X68,X70,X71),X73)|X72!=X73|~m1_subset_1(X73,u1_struct_0(X71))|~m1_subset_1(X72,u1_struct_0(X69))|(v2_struct_0(X71)|(~v1_tsep_1(X71,X69)|~m1_pre_topc(X71,X69)))|(~v1_funct_1(X70)|(~v1_funct_2(X70,u1_struct_0(X69),u1_struct_0(X68))|~m1_subset_1(X70,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X69),u1_struct_0(X68))))))|(v2_struct_0(X69)|(~v2_pre_topc(X69)|~l1_pre_topc(X69)))|(v2_struct_0(X68)|(~v2_pre_topc(X68)|~l1_pre_topc(X68))))&(~r1_tmap_1(X71,X68,k2_tmap_1(X69,X68,X70,X71),X73)|r1_tmap_1(X69,X68,X70,X72)|X72!=X73|~m1_subset_1(X73,u1_struct_0(X71))|~m1_subset_1(X72,u1_struct_0(X69))|(v2_struct_0(X71)|(~v1_tsep_1(X71,X69)|~m1_pre_topc(X71,X69)))|(~v1_funct_1(X70)|(~v1_funct_2(X70,u1_struct_0(X69),u1_struct_0(X68))|~m1_subset_1(X70,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X69),u1_struct_0(X68))))))|(v2_struct_0(X69)|(~v2_pre_topc(X69)|~l1_pre_topc(X69)))|(v2_struct_0(X68)|(~v2_pre_topc(X68)|~l1_pre_topc(X68))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t67_tmap_1])])])])])).
fof(c_0_100, plain, ![X61, X62, X63]:(v2_struct_0(X61)|~l1_pre_topc(X61)|(v2_struct_0(X62)|~m1_pre_topc(X62,X61)|(~m1_subset_1(X63,u1_struct_0(X62))|m1_subset_1(X63,u1_struct_0(X61))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t55_pre_topc])])])])).
cnf(c_0_101, plain, (v1_tsep_1(X1,X2)|~v3_pre_topc(X3,X2)|X3!=u1_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X1,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_102, negated_conjecture, (v3_pre_topc(u1_struct_0(esk3_0),esk4_0)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk4_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_38])])).
cnf(c_0_103, negated_conjecture, (m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk4_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_85]), c_0_38])])).
cnf(c_0_104, negated_conjecture, (k3_tmap_1(X1,esk2_0,esk4_0,X2,esk5_0)=k2_partfun1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),esk5_0,u1_struct_0(X2))|v2_struct_0(X1)|~m1_pre_topc(X2,esk4_0)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_91])]), c_0_92])).
cnf(c_0_105, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk4_0)|~r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk4_0))), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_106, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk4_0))), inference(spm,[status(thm)],[c_0_95, c_0_85])).
cnf(c_0_107, negated_conjecture, (k2_partfun1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),esk5_0,u1_struct_0(X1))=k2_tmap_1(esk4_0,esk2_0,esk5_0,X1)|~m1_pre_topc(X1,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_87]), c_0_88]), c_0_89]), c_0_90]), c_0_38]), c_0_91]), c_0_97])]), c_0_92]), c_0_98])).
cnf(c_0_108, plain, (r1_tmap_1(X3,X2,X4,X6)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|~r1_tmap_1(X1,X2,k2_tmap_1(X3,X2,X4,X1),X5)|X6!=X5|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X3))|~v1_tsep_1(X1,X3)|~m1_pre_topc(X1,X3)|~v1_funct_1(X4)|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_109, plain, (v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(X3,u1_struct_0(X1))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_110, plain, (v1_tsep_1(X1,X2)|~v3_pre_topc(u1_struct_0(X1),X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_101]), c_0_68])).
cnf(c_0_111, negated_conjecture, (v3_pre_topc(u1_struct_0(esk3_0),esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_102, c_0_103])])).
cnf(c_0_112, negated_conjecture, (k3_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk5_0)=k2_partfun1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),esk5_0,u1_struct_0(esk3_0))|v2_struct_0(X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_104, c_0_85])).
cnf(c_0_113, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_105, c_0_106])])).
cnf(c_0_114, negated_conjecture, (k2_partfun1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),esk5_0,u1_struct_0(esk4_0))=k2_tmap_1(esk4_0,esk2_0,esk5_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_59]), c_0_38])])).
cnf(c_0_115, plain, (r1_tmap_1(X1,X2,X3,X4)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X5)|~v1_tsep_1(X5,X1)|~r1_tmap_1(X5,X2,k2_tmap_1(X1,X2,X3,X5),X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,u1_struct_0(X5))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~m1_pre_topc(X5,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_108]), c_0_109])).
cnf(c_0_116, negated_conjecture, (v1_tsep_1(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_85]), c_0_38]), c_0_97])])).
cnf(c_0_117, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_118, negated_conjecture, (r1_tmap_1(esk3_0,esk2_0,k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk5_0),esk7_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_119, negated_conjecture, (esk6_0=esk7_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_120, negated_conjecture, (k3_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk5_0)=k2_tmap_1(esk4_0,esk2_0,esk5_0,esk4_0)|v2_struct_0(X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_112, c_0_113]), c_0_114])).
cnf(c_0_121, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_122, negated_conjecture, (r1_tmap_1(esk4_0,X1,X2,X3)|v2_struct_0(X1)|~r1_tmap_1(esk3_0,X1,k2_tmap_1(esk4_0,X1,X2,esk3_0),X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(esk4_0))|~v1_funct_2(X2,u1_struct_0(esk4_0),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_85]), c_0_38]), c_0_97])]), c_0_117]), c_0_98]), c_0_113])).
cnf(c_0_123, negated_conjecture, (k2_tmap_1(esk4_0,esk2_0,esk5_0,esk3_0)=k2_tmap_1(esk4_0,esk2_0,esk5_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_85]), c_0_113]), c_0_114])).
cnf(c_0_124, negated_conjecture, (r1_tmap_1(esk3_0,esk2_0,k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk5_0),esk6_0)), inference(rw,[status(thm)],[c_0_118, c_0_119])).
cnf(c_0_125, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk5_0)=k2_tmap_1(esk4_0,esk2_0,esk5_0,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_33]), c_0_31]), c_0_51])]), c_0_121])).
cnf(c_0_126, negated_conjecture, (r1_tmap_1(esk4_0,esk2_0,esk5_0,X1)|~r1_tmap_1(esk3_0,esk2_0,k2_tmap_1(esk4_0,esk2_0,esk5_0,esk4_0),X1)|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_87]), c_0_123]), c_0_88]), c_0_89]), c_0_90]), c_0_91])]), c_0_92])).
cnf(c_0_127, negated_conjecture, (r1_tmap_1(esk3_0,esk2_0,k2_tmap_1(esk4_0,esk2_0,esk5_0,esk4_0),esk6_0)), inference(rw,[status(thm)],[c_0_124, c_0_125])).
cnf(c_0_128, negated_conjecture, (m1_subset_1(esk6_0,u1_struct_0(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_129, negated_conjecture, (~r1_tmap_1(esk4_0,esk2_0,esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_130, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_127]), c_0_128])]), c_0_129]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 131
# Proof object clause steps            : 86
# Proof object formula steps           : 45
# Proof object conjectures             : 59
# Proof object clause conjectures      : 56
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 41
# Proof object initial formulas used   : 22
# Proof object generating inferences   : 32
# Proof object simplifying inferences  : 101
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 68
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 67
# Processed clauses                    : 286
# ...of these trivial                  : 8
# ...subsumed                          : 16
# ...remaining for further processing  : 262
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 42
# Generated clauses                    : 390
# ...of the previous two non-trivial   : 220
# Contextual simplify-reflections      : 7
# Paramodulations                      : 382
# Factorizations                       : 0
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 148
#    Positive orientable unit clauses  : 55
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 88
# Current number of unprocessed clauses: 54
# ...number of literals in the above   : 235
# Current number of archived formulas  : 0
# Current number of archived clauses   : 108
# Clause-clause subsumption calls (NU) : 2881
# Rec. Clause-clause subsumption calls : 261
# Non-unit clause-clause subsumptions  : 24
# Unit Clause-clause subsumption calls : 238
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 26
# BW rewrite match successes           : 10
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 16465

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 28796 pages
