# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t39_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2))), file('data/MPTP/proofs/099c4bf5e02b4b2c85a4d1cd577cdbbc.E_input', t39_yellow_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/099c4bf5e02b4b2c85a4d1cd577cdbbc.E_input', t24_orders_2)).
fof(t19_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))))), file('data/MPTP/proofs/099c4bf5e02b4b2c85a4d1cd577cdbbc.E_input', t19_yellow_0)).
fof(d10_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_yellow_0(X1,X2)=>(X3=k2_yellow_0(X1,X2)<=>(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))))))), file('data/MPTP/proofs/099c4bf5e02b4b2c85a4d1cd577cdbbc.E_input', d10_yellow_0)).
fof(dt_k2_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/099c4bf5e02b4b2c85a4d1cd577cdbbc.E_input', dt_k2_yellow_0)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/099c4bf5e02b4b2c85a4d1cd577cdbbc.E_input', d8_lattice3)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/099c4bf5e02b4b2c85a4d1cd577cdbbc.E_input', d2_tarski)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/099c4bf5e02b4b2c85a4d1cd577cdbbc.E_input', redefinition_k6_domain_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/099c4bf5e02b4b2c85a4d1cd577cdbbc.E_input', t69_enumset1)).
fof(t7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2)))))))), file('data/MPTP/proofs/099c4bf5e02b4b2c85a4d1cd577cdbbc.E_input', t7_yellow_0)).
fof(t8_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4)))&(((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4))=>r1_lattice3(X1,k2_tarski(X3,X4),X2))&((r2_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2)))&((r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2))=>r2_lattice3(X1,k2_tarski(X3,X4),X2))))))))), file('data/MPTP/proofs/099c4bf5e02b4b2c85a4d1cd577cdbbc.E_input', t8_yellow_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/099c4bf5e02b4b2c85a4d1cd577cdbbc.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/099c4bf5e02b4b2c85a4d1cd577cdbbc.E_input', dt_l1_orders_2)).
fof(t30_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4)))))&((r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))=>(X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3)))))), file('data/MPTP/proofs/099c4bf5e02b4b2c85a4d1cd577cdbbc.E_input', t30_yellow_0)).
fof(t25_orders_2, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X3,X2))=>X2=X3)))), file('data/MPTP/proofs/099c4bf5e02b4b2c85a4d1cd577cdbbc.E_input', t25_orders_2)).
fof(c_0_15, plain, ![X3, X2, X1]:(epred2_3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))), introduced(definition)).
fof(c_0_16, plain, ![X3, X2, X1]:(epred2_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))), inference(split_equiv,[status(thm)],[c_0_15])).
fof(c_0_17, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2)))), inference(assume_negation,[status(cth)],[t39_yellow_0])).
fof(c_0_18, plain, ![X59, X60, X61, X62, X63]:(((r1_orders_2(X61,X62,X60)|(X62!=k11_lattice3(X61,X60,X59)|~r2_yellow_0(X61,k2_tarski(X60,X59)))|~m1_subset_1(X62,u1_struct_0(X61))|~epred2_3(X61,X60,X59))&((r1_orders_2(X61,X62,X59)|(X62!=k11_lattice3(X61,X60,X59)|~r2_yellow_0(X61,k2_tarski(X60,X59)))|~m1_subset_1(X62,u1_struct_0(X61))|~epred2_3(X61,X60,X59))&(~m1_subset_1(X63,u1_struct_0(X61))|(~r1_orders_2(X61,X63,X60)|~r1_orders_2(X61,X63,X59)|r1_orders_2(X61,X63,X62))|(X62!=k11_lattice3(X61,X60,X59)|~r2_yellow_0(X61,k2_tarski(X60,X59)))|~m1_subset_1(X62,u1_struct_0(X61))|~epred2_3(X61,X60,X59))))&(((X62=k11_lattice3(X61,X60,X59)|(m1_subset_1(esk7_4(X59,X60,X61,X62),u1_struct_0(X61))|~r1_orders_2(X61,X62,X59)|~r1_orders_2(X61,X62,X60))|~m1_subset_1(X62,u1_struct_0(X61))|~epred2_3(X61,X60,X59))&(r2_yellow_0(X61,k2_tarski(X60,X59))|(m1_subset_1(esk7_4(X59,X60,X61,X62),u1_struct_0(X61))|~r1_orders_2(X61,X62,X59)|~r1_orders_2(X61,X62,X60))|~m1_subset_1(X62,u1_struct_0(X61))|~epred2_3(X61,X60,X59)))&((((X62=k11_lattice3(X61,X60,X59)|(r1_orders_2(X61,esk7_4(X59,X60,X61,X62),X60)|~r1_orders_2(X61,X62,X59)|~r1_orders_2(X61,X62,X60))|~m1_subset_1(X62,u1_struct_0(X61))|~epred2_3(X61,X60,X59))&(r2_yellow_0(X61,k2_tarski(X60,X59))|(r1_orders_2(X61,esk7_4(X59,X60,X61,X62),X60)|~r1_orders_2(X61,X62,X59)|~r1_orders_2(X61,X62,X60))|~m1_subset_1(X62,u1_struct_0(X61))|~epred2_3(X61,X60,X59)))&((X62=k11_lattice3(X61,X60,X59)|(r1_orders_2(X61,esk7_4(X59,X60,X61,X62),X59)|~r1_orders_2(X61,X62,X59)|~r1_orders_2(X61,X62,X60))|~m1_subset_1(X62,u1_struct_0(X61))|~epred2_3(X61,X60,X59))&(r2_yellow_0(X61,k2_tarski(X60,X59))|(r1_orders_2(X61,esk7_4(X59,X60,X61,X62),X59)|~r1_orders_2(X61,X62,X59)|~r1_orders_2(X61,X62,X60))|~m1_subset_1(X62,u1_struct_0(X61))|~epred2_3(X61,X60,X59))))&((X62=k11_lattice3(X61,X60,X59)|(~r1_orders_2(X61,esk7_4(X59,X60,X61,X62),X62)|~r1_orders_2(X61,X62,X59)|~r1_orders_2(X61,X62,X60))|~m1_subset_1(X62,u1_struct_0(X61))|~epred2_3(X61,X60,X59))&(r2_yellow_0(X61,k2_tarski(X60,X59))|(~r1_orders_2(X61,esk7_4(X59,X60,X61,X62),X62)|~r1_orders_2(X61,X62,X59)|~r1_orders_2(X61,X62,X60))|~m1_subset_1(X62,u1_struct_0(X61))|~epred2_3(X61,X60,X59)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])])).
fof(c_0_19, plain, ![X27, X28]:(v2_struct_0(X27)|(~v3_orders_2(X27)|~l1_orders_2(X27))|(~m1_subset_1(X28,u1_struct_0(X27))|r1_orders_2(X27,X28,X28))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_20, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])).
cnf(c_0_21, plain, (X1=k11_lattice3(X2,X3,X4)|~r1_orders_2(X2,esk7_4(X4,X3,X2,X1),X1)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~epred2_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, plain, (X1=k11_lattice3(X2,X3,X4)|r1_orders_2(X2,esk7_4(X4,X3,X2,X1),X4)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~epred2_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (X1=k11_lattice3(X2,X3,X4)|r1_orders_2(X2,esk7_4(X4,X3,X2,X1),X3)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~epred2_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, plain, (k11_lattice3(X1,X2,X3)=X3|~epred2_3(X1,X2,X3)|~r1_orders_2(X1,X3,X3)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_29, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_30, plain, (k11_lattice3(X1,X2,X3)=X2|~epred2_3(X1,X2,X3)|~r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_21, c_0_27])).
fof(c_0_31, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>epred2_3(X1,X2,X3)))), inference(apply_def,[status(thm)],[t19_yellow_0, c_0_15])).
fof(c_0_32, plain, ![X8, X9, X10, X11]:(((r1_lattice3(X8,X9,X10)|X10!=k2_yellow_0(X8,X9)|~r2_yellow_0(X8,X9)|~m1_subset_1(X10,u1_struct_0(X8))|~l1_orders_2(X8))&(~m1_subset_1(X11,u1_struct_0(X8))|(~r1_lattice3(X8,X9,X11)|r1_orders_2(X8,X11,X10))|X10!=k2_yellow_0(X8,X9)|~r2_yellow_0(X8,X9)|~m1_subset_1(X10,u1_struct_0(X8))|~l1_orders_2(X8)))&((m1_subset_1(esk3_3(X8,X9,X10),u1_struct_0(X8))|~r1_lattice3(X8,X9,X10)|X10=k2_yellow_0(X8,X9)|~r2_yellow_0(X8,X9)|~m1_subset_1(X10,u1_struct_0(X8))|~l1_orders_2(X8))&((r1_lattice3(X8,X9,esk3_3(X8,X9,X10))|~r1_lattice3(X8,X9,X10)|X10=k2_yellow_0(X8,X9)|~r2_yellow_0(X8,X9)|~m1_subset_1(X10,u1_struct_0(X8))|~l1_orders_2(X8))&(~r1_orders_2(X8,esk3_3(X8,X9,X10),X10)|~r1_lattice3(X8,X9,X10)|X10=k2_yellow_0(X8,X9)|~r2_yellow_0(X8,X9)|~m1_subset_1(X10,u1_struct_0(X8))|~l1_orders_2(X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_yellow_0])])])])])).
fof(c_0_33, plain, ![X25, X26]:(~l1_orders_2(X25)|m1_subset_1(k2_yellow_0(X25,X26),u1_struct_0(X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_yellow_0])])).
cnf(c_0_34, negated_conjecture, (k11_lattice3(esk1_0,X1,X2)=X2|~epred2_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35, negated_conjecture, (k11_lattice3(esk1_0,X1,X2)=X1|~epred2_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_30, c_0_29])).
fof(c_0_36, plain, ![X38, X39, X40]:(~v5_orders_2(X38)|~l1_orders_2(X38)|(~m1_subset_1(X39,u1_struct_0(X38))|(~m1_subset_1(X40,u1_struct_0(X38))|epred2_3(X38,X39,X40)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])])).
fof(c_0_37, plain, ![X29, X30, X31, X32]:((~r1_lattice3(X29,X30,X31)|(~m1_subset_1(X32,u1_struct_0(X29))|(~r2_hidden(X32,X30)|r1_orders_2(X29,X31,X32)))|~m1_subset_1(X31,u1_struct_0(X29))|~l1_orders_2(X29))&((m1_subset_1(esk5_3(X29,X30,X31),u1_struct_0(X29))|r1_lattice3(X29,X30,X31)|~m1_subset_1(X31,u1_struct_0(X29))|~l1_orders_2(X29))&((r2_hidden(esk5_3(X29,X30,X31),X30)|r1_lattice3(X29,X30,X31)|~m1_subset_1(X31,u1_struct_0(X29))|~l1_orders_2(X29))&(~r1_orders_2(X29,X31,esk5_3(X29,X30,X31))|r1_lattice3(X29,X30,X31)|~m1_subset_1(X31,u1_struct_0(X29))|~l1_orders_2(X29))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
cnf(c_0_38, plain, (r1_lattice3(X1,X2,X3)|X3!=k2_yellow_0(X1,X2)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39, plain, (m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_40, plain, ![X41, X42, X43, X44, X45, X46, X47, X48]:(((~r2_hidden(X44,X43)|(X44=X41|X44=X42)|X43!=k2_tarski(X41,X42))&((X45!=X41|r2_hidden(X45,X43)|X43!=k2_tarski(X41,X42))&(X45!=X42|r2_hidden(X45,X43)|X43!=k2_tarski(X41,X42))))&(((esk6_3(X46,X47,X48)!=X46|~r2_hidden(esk6_3(X46,X47,X48),X48)|X48=k2_tarski(X46,X47))&(esk6_3(X46,X47,X48)!=X47|~r2_hidden(esk6_3(X46,X47,X48),X48)|X48=k2_tarski(X46,X47)))&(r2_hidden(esk6_3(X46,X47,X48),X48)|(esk6_3(X46,X47,X48)=X46|esk6_3(X46,X47,X48)=X47)|X48=k2_tarski(X46,X47)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_41, negated_conjecture, (X1=X2|~epred2_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,X1)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_42, plain, (epred2_3(X1,X2,X3)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_43, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_44, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_45, plain, (r1_lattice3(X1,X2,k2_yellow_0(X1,X2))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_38]), c_0_39])).
cnf(c_0_46, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_47, negated_conjecture, (X1=X2|~r1_orders_2(esk1_0,X2,X1)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_25]), c_0_43])])).
cnf(c_0_48, plain, (r1_orders_2(X1,k2_yellow_0(X1,X2),X3)|~r2_hidden(X3,X2)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_39])).
cnf(c_0_49, plain, (r1_orders_2(X2,X1,X4)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|X4!=k2_yellow_0(X2,X3)|~r2_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_50, plain, ![X53, X54]:(v1_xboole_0(X53)|~m1_subset_1(X54,X53)|k6_domain_1(X53,X54)=k1_tarski(X54)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_51, plain, ![X24]:k2_tarski(X24,X24)=k1_tarski(X24), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_52, plain, (X1=X2|X1=X3|~r2_hidden(X1,k2_tarski(X3,X2))), inference(er,[status(thm)],[c_0_46])).
cnf(c_0_53, plain, (r2_hidden(esk5_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_54, negated_conjecture, (X1=k2_yellow_0(esk1_0,X2)|~r2_hidden(X1,X2)|~r1_orders_2(esk1_0,X1,k2_yellow_0(esk1_0,X2))|~r2_yellow_0(esk1_0,X2)|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_25])])).
cnf(c_0_55, plain, (r1_orders_2(X1,X2,k2_yellow_0(X1,X3))|~r1_lattice3(X1,X3,X2)|~r2_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_49]), c_0_39])).
cnf(c_0_56, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_57, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_58, plain, (esk5_3(X1,k2_tarski(X2,X3),X4)=X3|esk5_3(X1,k2_tarski(X2,X3),X4)=X2|r1_lattice3(X1,k2_tarski(X2,X3),X4)|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_59, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_60, negated_conjecture, (X1=k2_yellow_0(esk1_0,X2)|~r2_hidden(X1,X2)|~r1_lattice3(esk1_0,X2,X1)|~r2_yellow_0(esk1_0,X2)|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_25])])).
cnf(c_0_61, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_62, plain, ![X50, X51, X52]:((~r1_lattice3(X50,k1_tarski(X52),X51)|r1_orders_2(X50,X51,X52)|~m1_subset_1(X52,u1_struct_0(X50))|~m1_subset_1(X51,u1_struct_0(X50))|~l1_orders_2(X50))&((~r1_orders_2(X50,X51,X52)|r1_lattice3(X50,k1_tarski(X52),X51)|~m1_subset_1(X52,u1_struct_0(X50))|~m1_subset_1(X51,u1_struct_0(X50))|~l1_orders_2(X50))&((~r2_lattice3(X50,k1_tarski(X52),X51)|r1_orders_2(X50,X52,X51)|~m1_subset_1(X52,u1_struct_0(X50))|~m1_subset_1(X51,u1_struct_0(X50))|~l1_orders_2(X50))&(~r1_orders_2(X50,X52,X51)|r2_lattice3(X50,k1_tarski(X52),X51)|~m1_subset_1(X52,u1_struct_0(X50))|~m1_subset_1(X51,u1_struct_0(X50))|~l1_orders_2(X50))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])).
fof(c_0_63, plain, ![X4, X3, X2, X1]:(epred1_4(X1,X2,X3,X4)<=>((r1_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4)))&(((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4))=>r1_lattice3(X1,k2_tarski(X3,X4),X2))&((r2_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2)))&((r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2))=>r2_lattice3(X1,k2_tarski(X3,X4),X2)))))), introduced(definition)).
cnf(c_0_64, plain, (k2_tarski(X2,X2)=k6_domain_1(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_65, plain, (m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_66, negated_conjecture, (esk5_3(esk1_0,k2_tarski(X1,X2),esk2_0)=X1|esk5_3(esk1_0,k2_tarski(X1,X2),esk2_0)=X2|r1_lattice3(esk1_0,k2_tarski(X1,X2),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_25])])).
cnf(c_0_67, negated_conjecture, (X1=k2_yellow_0(esk1_0,X2)|~r2_hidden(X1,X2)|~r1_lattice3(esk1_0,X2,X1)|~r2_yellow_0(esk1_0,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_39]), c_0_25])])).
cnf(c_0_68, plain, (r2_hidden(X1,k2_tarski(X1,X2))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_61])])).
cnf(c_0_69, plain, (r1_lattice3(X1,k1_tarski(X3),X2)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
fof(c_0_70, plain, ![X4, X3, X2, X1]:(epred1_4(X1,X2,X3,X4)=>((r1_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4)))&(((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4))=>r1_lattice3(X1,k2_tarski(X3,X4),X2))&((r2_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2)))&((r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2))=>r2_lattice3(X1,k2_tarski(X3,X4),X2)))))), inference(split_equiv,[status(thm)],[c_0_63])).
fof(c_0_71, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>epred1_4(X1,X2,X3,X4))))), inference(apply_def,[status(thm)],[t8_yellow_0, c_0_63])).
fof(c_0_72, plain, ![X34]:(v2_struct_0(X34)|~l1_struct_0(X34)|~v1_xboole_0(u1_struct_0(X34))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_73, plain, (k2_tarski(esk5_3(X1,X2,X3),esk5_3(X1,X2,X3))=k6_domain_1(u1_struct_0(X1),esk5_3(X1,X2,X3))|v1_xboole_0(u1_struct_0(X1))|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_74, negated_conjecture, (esk5_3(esk1_0,k2_tarski(X1,X1),esk2_0)=X1|r1_lattice3(esk1_0,k2_tarski(X1,X1),esk2_0)), inference(er,[status(thm)],[inference(ef,[status(thm)],[c_0_66])])).
cnf(c_0_75, negated_conjecture, (k2_yellow_0(esk1_0,k2_tarski(X1,X2))=X1|~r1_lattice3(esk1_0,k2_tarski(X1,X2),X1)|~r2_yellow_0(esk1_0,k2_tarski(X1,X2))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_76, plain, (r1_lattice3(X1,k2_tarski(X3,X3),X2)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[c_0_69, c_0_57])).
cnf(c_0_77, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|~r1_orders_2(X1,esk7_4(X3,X2,X1,X4),X4)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~epred2_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_78, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|r1_orders_2(X1,esk7_4(X3,X2,X1,X4),X3)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~epred2_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_79, plain, ![X55, X56, X57, X58]:(((r1_orders_2(X58,X57,X56)|~r1_lattice3(X58,k2_tarski(X56,X55),X57)|~epred1_4(X58,X57,X56,X55))&(r1_orders_2(X58,X57,X55)|~r1_lattice3(X58,k2_tarski(X56,X55),X57)|~epred1_4(X58,X57,X56,X55)))&((~r1_orders_2(X58,X57,X56)|~r1_orders_2(X58,X57,X55)|r1_lattice3(X58,k2_tarski(X56,X55),X57)|~epred1_4(X58,X57,X56,X55))&(((r1_orders_2(X58,X56,X57)|~r2_lattice3(X58,k2_tarski(X56,X55),X57)|~epred1_4(X58,X57,X56,X55))&(r1_orders_2(X58,X55,X57)|~r2_lattice3(X58,k2_tarski(X56,X55),X57)|~epred1_4(X58,X57,X56,X55)))&(~r1_orders_2(X58,X56,X57)|~r1_orders_2(X58,X55,X57)|r2_lattice3(X58,k2_tarski(X56,X55),X57)|~epred1_4(X58,X57,X56,X55))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_70])])])).
fof(c_0_80, plain, ![X19, X20, X21, X22]:(~l1_orders_2(X19)|(~m1_subset_1(X20,u1_struct_0(X19))|(~m1_subset_1(X21,u1_struct_0(X19))|(~m1_subset_1(X22,u1_struct_0(X19))|epred1_4(X19,X20,X21,X22))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_71])])])).
cnf(c_0_81, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_82, negated_conjecture, (k2_tarski(X1,X1)=k6_domain_1(u1_struct_0(esk1_0),X1)|v1_xboole_0(u1_struct_0(esk1_0))|r1_lattice3(esk1_0,k2_tarski(X1,X1),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_59]), c_0_25])])).
fof(c_0_83, plain, ![X23]:(~l1_orders_2(X23)|l1_struct_0(X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_84, negated_conjecture, (k2_yellow_0(esk1_0,k2_tarski(X1,X1))=X1|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_25])]), c_0_29])).
cnf(c_0_85, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|~epred2_3(X1,X2,X3)|~r1_orders_2(X1,X3,X3)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_86, plain, (r1_orders_2(X1,X2,X3)|~r1_lattice3(X1,k2_tarski(X3,X4),X2)|~epred1_4(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_87, plain, (epred1_4(X1,X2,X3,X4)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_88, negated_conjecture, (k2_tarski(X1,X1)=k6_domain_1(u1_struct_0(esk1_0),X1)|r1_lattice3(esk1_0,k2_tarski(X1,X1),esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_26])).
cnf(c_0_89, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_90, negated_conjecture, (r1_lattice3(esk1_0,k2_tarski(X1,X1),X1)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_84]), c_0_25])])).
cnf(c_0_91, negated_conjecture, (r2_yellow_0(esk1_0,k2_tarski(X1,X2))|~epred2_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_85, c_0_29])).
cnf(c_0_92, plain, (r1_orders_2(X1,X2,X3)|~r1_lattice3(X1,k2_tarski(X3,X4),X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_93, negated_conjecture, (k2_tarski(X1,X1)=k6_domain_1(u1_struct_0(esk1_0),X1)|r1_lattice3(esk1_0,k2_tarski(X1,X1),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_25])])).
cnf(c_0_94, negated_conjecture, (r1_lattice3(esk1_0,k2_tarski(X1,X1),X1)|~epred2_3(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_29])).
cnf(c_0_95, plain, (k2_tarski(X1,X1)=k6_domain_1(u1_struct_0(esk1_0),X1)|r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_59]), c_0_25])])).
cnf(c_0_96, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_64, c_0_59])).
cnf(c_0_97, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk5_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_98, plain, (r1_orders_2(esk1_0,esk2_0,X1)|r1_lattice3(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1),X1)|~epred2_3(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_99, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_96]), c_0_26])).
cnf(c_0_100, negated_conjecture, (r1_lattice3(esk1_0,k2_tarski(X1,X1),esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_74]), c_0_59]), c_0_25])])).
cnf(c_0_101, negated_conjecture, (r1_lattice3(esk1_0,k2_tarski(esk2_0,esk2_0),esk2_0)|~epred2_3(esk1_0,esk2_0,esk2_0)|~l1_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_59])]), c_0_100])).
fof(c_0_102, plain, ![X13, X14, X15, X16, X17]:(((r2_lattice3(X13,X15,X14)|(X14!=k1_yellow_0(X13,X15)|~r1_yellow_0(X13,X15))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13)))&(~m1_subset_1(X16,u1_struct_0(X13))|(~r2_lattice3(X13,X15,X16)|r1_orders_2(X13,X14,X16))|(X14!=k1_yellow_0(X13,X15)|~r1_yellow_0(X13,X15))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13))))&(((X14=k1_yellow_0(X13,X17)|(m1_subset_1(esk4_3(X13,X14,X17),u1_struct_0(X13))|~r2_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13)))&(r1_yellow_0(X13,X17)|(m1_subset_1(esk4_3(X13,X14,X17),u1_struct_0(X13))|~r2_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13))))&(((X14=k1_yellow_0(X13,X17)|(r2_lattice3(X13,X17,esk4_3(X13,X14,X17))|~r2_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13)))&(r1_yellow_0(X13,X17)|(r2_lattice3(X13,X17,esk4_3(X13,X14,X17))|~r2_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13))))&((X14=k1_yellow_0(X13,X17)|(~r1_orders_2(X13,X14,esk4_3(X13,X14,X17))|~r2_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13)))&(r1_yellow_0(X13,X17)|(~r1_orders_2(X13,X14,esk4_3(X13,X14,X17))|~r2_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_yellow_0])])])])])])).
cnf(c_0_103, negated_conjecture, (esk5_3(esk1_0,k2_tarski(X1,X2),esk2_0)=X1|r1_lattice3(esk1_0,k2_tarski(X1,X2),esk2_0)|~r1_orders_2(esk1_0,esk2_0,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_66]), c_0_59]), c_0_25])])).
cnf(c_0_104, negated_conjecture, (r1_lattice3(esk1_0,k2_tarski(esk2_0,esk2_0),esk2_0)|~epred2_3(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_89]), c_0_25])])).
cnf(c_0_105, plain, (X1=k1_yellow_0(X2,X3)|~r1_orders_2(X2,X1,esk4_3(X2,X1,X3))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_106, negated_conjecture, (r1_lattice3(esk1_0,k2_tarski(X1,X2),esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,esk2_0,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_103]), c_0_59]), c_0_25])])).
cnf(c_0_107, plain, (r1_orders_2(esk1_0,esk2_0,esk2_0)|~epred2_3(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_104]), c_0_59]), c_0_25])])).
cnf(c_0_108, plain, (k2_yellow_0(X1,X2)=k1_yellow_0(X1,X3)|~r2_hidden(esk4_3(X1,k2_yellow_0(X1,X2),X3),X2)|~r2_lattice3(X1,X3,k2_yellow_0(X1,X2))|~r2_yellow_0(X1,X2)|~m1_subset_1(esk4_3(X1,k2_yellow_0(X1,X2),X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_48]), c_0_39])).
cnf(c_0_109, plain, (X1=k1_yellow_0(X2,X3)|m1_subset_1(esk4_3(X2,X1,X3),u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_110, negated_conjecture, (k2_yellow_0(esk1_0,k2_tarski(esk2_0,X1))=esk2_0|~r1_orders_2(esk1_0,esk2_0,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_106]), c_0_59])])).
cnf(c_0_111, plain, (r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_42]), c_0_59]), c_0_25]), c_0_43])])).
cnf(c_0_112, plain, (k2_yellow_0(X1,X2)=k1_yellow_0(X1,X3)|~r2_hidden(esk4_3(X1,k2_yellow_0(X1,X2),X3),X2)|~r2_lattice3(X1,X3,k2_yellow_0(X1,X2))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_39])).
cnf(c_0_113, negated_conjecture, (k2_yellow_0(esk1_0,k2_tarski(esk2_0,X1))=esk2_0|~r1_orders_2(esk1_0,esk2_0,X1)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_110, c_0_111])])).
cnf(c_0_114, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_115, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|r1_orders_2(X1,esk7_4(X3,X2,X1,X4),X2)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~epred2_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_116, negated_conjecture, (k1_yellow_0(esk1_0,X1)=esk2_0|~r2_hidden(esk4_3(esk1_0,esk2_0,X1),k2_tarski(esk2_0,X2))|~r2_lattice3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X2)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_25]), c_0_43])])).
cnf(c_0_117, plain, (r2_hidden(X1,k2_tarski(X2,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_114])])).
cnf(c_0_118, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|~epred2_3(X1,X2,X3)|~r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_77, c_0_115])).
cnf(c_0_119, negated_conjecture, (k1_yellow_0(esk1_0,X1)=esk2_0|~r2_lattice3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,esk4_3(esk1_0,esk2_0,X1))|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_3(esk1_0,esk2_0,X1)))), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_120, negated_conjecture, (r2_yellow_0(esk1_0,k2_tarski(X1,X2))|~epred2_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_118, c_0_29])).
cnf(c_0_121, plain, (r1_orders_2(X1,X2,X3)|~r2_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_122, plain, (r1_lattice3(X1,k2_tarski(X3,X4),X2)|~r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,X2,X4)|~epred1_4(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_123, negated_conjecture, (k1_yellow_0(esk1_0,X1)=esk2_0|~epred2_3(esk1_0,esk2_0,esk4_3(esk1_0,esk2_0,X1))|~r2_lattice3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,esk4_3(esk1_0,esk2_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120]), c_0_59])])).
cnf(c_0_124, plain, (r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r2_lattice3(X1,k2_tarski(X2,X2),X3)), inference(rw,[status(thm)],[c_0_121, c_0_57])).
cnf(c_0_125, plain, (X1=k1_yellow_0(X2,X3)|r2_lattice3(X2,X3,esk4_3(X2,X1,X3))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_102])).
fof(c_0_126, plain, ![X35, X36, X37]:(~v5_orders_2(X35)|~l1_orders_2(X35)|(~m1_subset_1(X36,u1_struct_0(X35))|(~m1_subset_1(X37,u1_struct_0(X35))|(~r1_orders_2(X35,X36,X37)|~r1_orders_2(X35,X37,X36)|X36=X37)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t25_orders_2])])])).
cnf(c_0_127, plain, (r1_lattice3(X1,k2_tarski(X2,X3),X4)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_122, c_0_87])).
cnf(c_0_128, negated_conjecture, (k1_yellow_0(esk1_0,X1)=esk2_0|~r2_lattice3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,esk4_3(esk1_0,esk2_0,X1))|~m1_subset_1(esk4_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_42]), c_0_59]), c_0_25]), c_0_43])])).
cnf(c_0_129, plain, (X1=k1_yellow_0(X2,k2_tarski(X3,X3))|r1_orders_2(X2,X3,esk4_3(X2,X1,k2_tarski(X3,X3)))|~r2_lattice3(X2,k2_tarski(X3,X3),X1)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125]), c_0_109])).
cnf(c_0_130, plain, (X2=X3|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_126])).
cnf(c_0_131, plain, (k2_yellow_0(esk1_0,k2_tarski(X1,X2))=X1|~r1_orders_2(esk1_0,X1,X2)|~r2_yellow_0(esk1_0,k2_tarski(X1,X2))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_127]), c_0_25])]), c_0_29])).
cnf(c_0_132, negated_conjecture, (k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_133, negated_conjecture, (k1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))=esk2_0|~r2_lattice3(esk1_0,k2_tarski(esk2_0,esk2_0),esk2_0)|~m1_subset_1(esk4_3(esk1_0,esk2_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129]), c_0_59]), c_0_25]), c_0_43])])).
cnf(c_0_134, plain, (X1=k2_yellow_0(X2,X3)|~r2_hidden(X1,X3)|~r1_orders_2(X2,X1,k2_yellow_0(X2,X3))|~r2_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_48]), c_0_39])).
cnf(c_0_135, plain, (r1_lattice3(esk1_0,k2_tarski(X1,X2),X1)|~r1_orders_2(esk1_0,X1,X2)|~r2_yellow_0(esk1_0,k2_tarski(X1,X2))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_131]), c_0_25])])).
cnf(c_0_136, plain, (r2_yellow_0(esk1_0,k2_tarski(X1,esk2_0))|~epred2_3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_111]), c_0_59])])).
cnf(c_0_137, negated_conjecture, (k1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))!=esk2_0|~l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_132, c_0_99])).
cnf(c_0_138, negated_conjecture, (k1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))=esk2_0|~r2_lattice3(esk1_0,k2_tarski(esk2_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_109]), c_0_59]), c_0_25]), c_0_43])])).
cnf(c_0_139, plain, (X1=k2_yellow_0(X2,X3)|~r2_hidden(X1,X3)|~r1_lattice3(X2,X3,X1)|~r2_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)), inference(spm,[status(thm)],[c_0_134, c_0_55])).
cnf(c_0_140, plain, (r1_lattice3(esk1_0,k2_tarski(X1,esk2_0),X1)|~epred2_3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_136]), c_0_59])])).
cnf(c_0_141, negated_conjecture, (k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))!=esk2_0|~l1_struct_0(esk1_0)|~r2_lattice3(esk1_0,k2_tarski(esk2_0,esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_137, c_0_138])).
cnf(c_0_142, plain, (k2_yellow_0(esk1_0,k2_tarski(X1,esk2_0))=X1|~epred2_3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140]), c_0_68]), c_0_25]), c_0_43])]), c_0_136])).
cnf(c_0_143, plain, (~epred2_3(esk1_0,esk2_0,esk2_0)|~l1_struct_0(esk1_0)|~r2_lattice3(esk1_0,k2_tarski(esk2_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_142]), c_0_111]), c_0_59])])).
cnf(c_0_144, plain, (r2_lattice3(X1,k2_tarski(X2,X4),X3)|~r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,X4,X3)|~epred1_4(X1,X3,X2,X4)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_145, plain, (~epred2_3(esk1_0,esk2_0,esk2_0)|~r2_lattice3(esk1_0,k2_tarski(esk2_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_89]), c_0_25])])).
cnf(c_0_146, plain, (r2_lattice3(X1,k2_tarski(X2,X3),X4)|~r1_orders_2(X1,X3,X4)|~r1_orders_2(X1,X2,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_144, c_0_87])).
cnf(c_0_147, plain, (~epred2_3(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146]), c_0_111]), c_0_59]), c_0_25])])).
cnf(c_0_148, plain, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_42]), c_0_59]), c_0_25]), c_0_43])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 149
# Proof object clause steps            : 110
# Proof object formula steps           : 39
# Proof object conjectures             : 45
# Proof object clause conjectures      : 42
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 38
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 63
# Proof object simplifying inferences  : 117
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 59
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 58
# Processed clauses                    : 7051
# ...of these trivial                  : 33
# ...subsumed                          : 5368
# ...remaining for further processing  : 1650
# Other redundant clauses eliminated   : 122
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 368
# Backward-rewritten                   : 49
# Generated clauses                    : 42312
# ...of the previous two non-trivial   : 35219
# Contextual simplify-reflections      : 272
# Paramodulations                      : 42114
# Factorizations                       : 78
# Equation resolutions                 : 122
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1165
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 1155
# Current number of unprocessed clauses: 27183
# ...number of literals in the above   : 222818
# Current number of archived formulas  : 0
# Current number of archived clauses   : 476
# Clause-clause subsumption calls (NU) : 195370
# Rec. Clause-clause subsumption calls : 23523
# Non-unit clause-clause subsumptions  : 6007
# Unit Clause-clause subsumption calls : 440
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1439277

# -------------------------------------------------
# User time                : 0.801 s
# System time              : 0.007 s
# Total time               : 0.808 s
# Maximum resident set size: 30720 pages
