# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/0a3cb8c35146499b8b3fa40dce4b5b9c.E_input', t36_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/0a3cb8c35146499b8b3fa40dce4b5b9c.E_input', t3_boole)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/0a3cb8c35146499b8b3fa40dce4b5b9c.E_input', l32_xboole_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/0a3cb8c35146499b8b3fa40dce4b5b9c.E_input', d5_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/0a3cb8c35146499b8b3fa40dce4b5b9c.E_input', d3_tarski)).
fof(d1_setfam_1, axiom, ![X1, X2]:((~(X1=k1_xboole_0)=>(X2=k1_setfam_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>![X4]:(r2_hidden(X4,X1)=>r2_hidden(X3,X4)))))&(X1=k1_xboole_0=>(X2=k1_setfam_1(X1)<=>X2=k1_xboole_0))), file('data/MPTP/proofs/0a3cb8c35146499b8b3fa40dce4b5b9c.E_input', d1_setfam_1)).
fof(t4_setfam_1, conjecture, ![X1, X2]:(r2_hidden(X1,X2)=>r1_tarski(k1_setfam_1(X2),X1)), file('data/MPTP/proofs/0a3cb8c35146499b8b3fa40dce4b5b9c.E_input', t4_setfam_1)).
fof(c_0_7, plain, ![X8, X9]:r1_tarski(k4_xboole_0(X8,X9),X8), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_8, plain, ![X7]:k4_xboole_0(X7,k1_xboole_0)=X7, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_9, plain, ![X31, X32]:((k4_xboole_0(X31,X32)!=k1_xboole_0|r1_tarski(X31,X32))&(~r1_tarski(X31,X32)|k4_xboole_0(X31,X32)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_10, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
fof(c_0_14, plain, ![X33, X34, X35, X36, X37, X38, X39, X40]:((((r2_hidden(X36,X33)|~r2_hidden(X36,X35)|X35!=k4_xboole_0(X33,X34))&(~r2_hidden(X36,X34)|~r2_hidden(X36,X35)|X35!=k4_xboole_0(X33,X34)))&(~r2_hidden(X37,X33)|r2_hidden(X37,X34)|r2_hidden(X37,X35)|X35!=k4_xboole_0(X33,X34)))&((~r2_hidden(esk7_3(X38,X39,X40),X40)|(~r2_hidden(esk7_3(X38,X39,X40),X38)|r2_hidden(esk7_3(X38,X39,X40),X39))|X40=k4_xboole_0(X38,X39))&((r2_hidden(esk7_3(X38,X39,X40),X38)|r2_hidden(esk7_3(X38,X39,X40),X40)|X40=k4_xboole_0(X38,X39))&(~r2_hidden(esk7_3(X38,X39,X40),X39)|r2_hidden(esk7_3(X38,X39,X40),X40)|X40=k4_xboole_0(X38,X39))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_15, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_12, c_0_13])).
fof(c_0_16, plain, ![X25, X26, X27, X28, X29]:((~r1_tarski(X25,X26)|(~r2_hidden(X27,X25)|r2_hidden(X27,X26)))&((r2_hidden(esk6_2(X28,X29),X28)|r1_tarski(X28,X29))&(~r2_hidden(esk6_2(X28,X29),X29)|r1_tarski(X28,X29)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_17, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_10, c_0_15])).
cnf(c_0_19, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_11])])).
fof(c_0_21, plain, ![X12, X13, X14, X15, X16, X18, X21, X22, X23, X24]:((((~r2_hidden(X14,X13)|(~r2_hidden(X15,X12)|r2_hidden(X14,X15))|X13!=k1_setfam_1(X12)|X12=k1_xboole_0)&((r2_hidden(esk3_3(X12,X13,X16),X12)|r2_hidden(X16,X13)|X13!=k1_setfam_1(X12)|X12=k1_xboole_0)&(~r2_hidden(X16,esk3_3(X12,X13,X16))|r2_hidden(X16,X13)|X13!=k1_setfam_1(X12)|X12=k1_xboole_0)))&(((r2_hidden(esk5_2(X12,X18),X12)|~r2_hidden(esk4_2(X12,X18),X18)|X18=k1_setfam_1(X12)|X12=k1_xboole_0)&(~r2_hidden(esk4_2(X12,X18),esk5_2(X12,X18))|~r2_hidden(esk4_2(X12,X18),X18)|X18=k1_setfam_1(X12)|X12=k1_xboole_0))&(r2_hidden(esk4_2(X12,X18),X18)|(~r2_hidden(X21,X12)|r2_hidden(esk4_2(X12,X18),X21))|X18=k1_setfam_1(X12)|X12=k1_xboole_0)))&((X23!=k1_setfam_1(X22)|X23=k1_xboole_0|X22!=k1_xboole_0)&(X24!=k1_xboole_0|X24=k1_setfam_1(X22)|X22!=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_setfam_1])])])])])])])).
cnf(c_0_22, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_12, c_0_18])).
cnf(c_0_24, plain, (~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_18]), c_0_20])).
fof(c_0_25, negated_conjecture, ~(![X1, X2]:(r2_hidden(X1,X2)=>r1_tarski(k1_setfam_1(X2),X1))), inference(assume_negation,[status(cth)],[t4_setfam_1])).
cnf(c_0_26, plain, (r2_hidden(X1,X3)|X4=k1_xboole_0|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)|X2!=k1_setfam_1(X4)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24])).
fof(c_0_28, negated_conjecture, (r2_hidden(esk1_0,esk2_0)&~r1_tarski(k1_setfam_1(esk2_0),esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])).
cnf(c_0_29, plain, (r2_hidden(X1,X2)|X3!=k1_setfam_1(X4)|~r2_hidden(X2,X4)|~r2_hidden(X1,X3)), inference(csr,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_30, negated_conjecture, (r2_hidden(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_31, negated_conjecture, (r2_hidden(X1,esk1_0)|X2!=k1_setfam_1(esk2_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_32, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X1,k1_setfam_1(esk2_0))), inference(er,[status(thm)],[c_0_31])).
cnf(c_0_33, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_34, plain, (r1_tarski(X1,X2)|~r2_hidden(esk6_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_35, negated_conjecture, (r1_tarski(k1_setfam_1(esk2_0),X1)|r2_hidden(esk6_2(k1_setfam_1(esk2_0),X1),esk1_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_36, negated_conjecture, (~r1_tarski(k1_setfam_1(esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 38
# Proof object clause steps            : 23
# Proof object formula steps           : 15
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 5
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 25
# Processed clauses                    : 90
# ...of these trivial                  : 2
# ...subsumed                          : 20
# ...remaining for further processing  : 68
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 2
# Generated clauses                    : 77
# ...of the previous two non-trivial   : 62
# Contextual simplify-reflections      : 2
# Paramodulations                      : 69
# Factorizations                       : 0
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 40
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 25
# Current number of unprocessed clauses: 19
# ...number of literals in the above   : 58
# Current number of archived formulas  : 0
# Current number of archived clauses   : 28
# Clause-clause subsumption calls (NU) : 93
# Rec. Clause-clause subsumption calls : 61
# Non-unit clause-clause subsumptions  : 11
# Unit Clause-clause subsumption calls : 37
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 11
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2217

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.000 s
# Total time               : 0.009 s
# Maximum resident set size: 28464 pages
