# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d11_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:((v1_pre_topc(X3)&(v2_pre_topc(X3)&l1_pre_topc(X3)))=>(X3=k8_tmap_1(X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(X4=u1_struct_0(X2)=>X3=k6_tmap_1(X1,X4))))))), file('data/MPTP/proofs/0a3e0424c8e1437db98b8004eac15bf4.E_input', d11_tmap_1)).
fof(dt_k8_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_pre_topc(X2,X1))=>(v1_pre_topc(k8_tmap_1(X1,X2))&(v2_pre_topc(k8_tmap_1(X1,X2))&l1_pre_topc(k8_tmap_1(X1,X2))))), file('data/MPTP/proofs/0a3e0424c8e1437db98b8004eac15bf4.E_input', dt_k8_tmap_1)).
fof(t1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/0a3e0424c8e1437db98b8004eac15bf4.E_input', t1_tsep_1)).
fof(t104_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(u1_struct_0(k6_tmap_1(X1,X2))=u1_struct_0(X1)&u1_pre_topc(k6_tmap_1(X1,X2))=k5_tmap_1(X1,X2)))), file('data/MPTP/proofs/0a3e0424c8e1437db98b8004eac15bf4.E_input', t104_tmap_1)).
fof(t121_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>(v1_funct_1(k2_tmap_1(X1,k8_tmap_1(X1,X2),k9_tmap_1(X1,X2),X2))&(v1_funct_2(k2_tmap_1(X1,k8_tmap_1(X1,X2),k9_tmap_1(X1,X2),X2),u1_struct_0(X2),u1_struct_0(k8_tmap_1(X1,X2)))&(v5_pre_topc(k2_tmap_1(X1,k8_tmap_1(X1,X2),k9_tmap_1(X1,X2),X2),X2,k8_tmap_1(X1,X2))&m1_subset_1(k2_tmap_1(X1,k8_tmap_1(X1,X2),k9_tmap_1(X1,X2),X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k8_tmap_1(X1,X2)))))))))), file('data/MPTP/proofs/0a3e0424c8e1437db98b8004eac15bf4.E_input', t121_tmap_1)).
fof(fc1_struct_0, axiom, ![X1]:((v2_struct_0(X1)&l1_struct_0(X1))=>v1_xboole_0(u1_struct_0(X1))), file('data/MPTP/proofs/0a3e0424c8e1437db98b8004eac15bf4.E_input', fc1_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/0a3e0424c8e1437db98b8004eac15bf4.E_input', dt_l1_pre_topc)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/0a3e0424c8e1437db98b8004eac15bf4.E_input', fc2_struct_0)).
fof(d10_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k7_tmap_1(X1,X2)=k6_partfun1(u1_struct_0(X1)))), file('data/MPTP/proofs/0a3e0424c8e1437db98b8004eac15bf4.E_input', d10_tmap_1)).
fof(dt_k7_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_funct_1(k7_tmap_1(X1,X2))&(v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))&m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))))), file('data/MPTP/proofs/0a3e0424c8e1437db98b8004eac15bf4.E_input', dt_k7_tmap_1)).
fof(d12_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))))=>(X3=k9_tmap_1(X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(X4=u1_struct_0(X2)=>r1_funct_2(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X4)),X3,k7_tmap_1(X1,X4)))))))), file('data/MPTP/proofs/0a3e0424c8e1437db98b8004eac15bf4.E_input', d12_tmap_1)).
fof(redefinition_r1_funct_2, axiom, ![X1, X2, X3, X4, X5, X6]:((~(v1_xboole_0(X2))&(~(v1_xboole_0(X4))&((v1_funct_1(X5)&(v1_funct_2(X5,X1,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&(v1_funct_1(X6)&(v1_funct_2(X6,X3,X4)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))))))=>(r1_funct_2(X1,X2,X3,X4,X5,X6)<=>X5=X6)), file('data/MPTP/proofs/0a3e0424c8e1437db98b8004eac15bf4.E_input', redefinition_r1_funct_2)).
fof(t112_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(u1_struct_0(X3)=X2=>(v1_funct_1(k2_tmap_1(X1,k6_tmap_1(X1,X2),k7_tmap_1(X1,X2),X3))&(v1_funct_2(k2_tmap_1(X1,k6_tmap_1(X1,X2),k7_tmap_1(X1,X2),X3),u1_struct_0(X3),u1_struct_0(k6_tmap_1(X1,X2)))&(v5_pre_topc(k2_tmap_1(X1,k6_tmap_1(X1,X2),k7_tmap_1(X1,X2),X3),X3,k6_tmap_1(X1,X2))&m1_subset_1(k2_tmap_1(X1,k6_tmap_1(X1,X2),k7_tmap_1(X1,X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k6_tmap_1(X1,X2)))))))))))), file('data/MPTP/proofs/0a3e0424c8e1437db98b8004eac15bf4.E_input', t112_tmap_1)).
fof(c_0_13, plain, ![X7, X8, X9, X10]:((X9!=k8_tmap_1(X7,X8)|(~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X7)))|(X10!=u1_struct_0(X8)|X9=k6_tmap_1(X7,X10)))|(~v1_pre_topc(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9)))|~m1_pre_topc(X8,X7)|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&((m1_subset_1(esk1_3(X7,X8,X9),k1_zfmisc_1(u1_struct_0(X7)))|X9=k8_tmap_1(X7,X8)|(~v1_pre_topc(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9)))|~m1_pre_topc(X8,X7)|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&((esk1_3(X7,X8,X9)=u1_struct_0(X8)|X9=k8_tmap_1(X7,X8)|(~v1_pre_topc(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9)))|~m1_pre_topc(X8,X7)|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&(X9!=k6_tmap_1(X7,esk1_3(X7,X8,X9))|X9=k8_tmap_1(X7,X8)|(~v1_pre_topc(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9)))|~m1_pre_topc(X8,X7)|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d11_tmap_1])])])])])])).
fof(c_0_14, plain, ![X12, X13]:((v1_pre_topc(k8_tmap_1(X12,X13))|(v2_struct_0(X12)|(~v2_pre_topc(X12)|~l1_pre_topc(X12))|~m1_pre_topc(X13,X12)))&((v2_pre_topc(k8_tmap_1(X12,X13))|(v2_struct_0(X12)|(~v2_pre_topc(X12)|~l1_pre_topc(X12))|~m1_pre_topc(X13,X12)))&(l1_pre_topc(k8_tmap_1(X12,X13))|(v2_struct_0(X12)|(~v2_pre_topc(X12)|~l1_pre_topc(X12))|~m1_pre_topc(X13,X12))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k8_tmap_1])])])])).
fof(c_0_15, plain, ![X14, X15]:(~l1_pre_topc(X14)|(~m1_pre_topc(X15,X14)|m1_subset_1(u1_struct_0(X15),k1_zfmisc_1(u1_struct_0(X14))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])).
fof(c_0_16, plain, ![X16, X17]:((u1_struct_0(k6_tmap_1(X16,X17))=u1_struct_0(X16)|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))))&(u1_pre_topc(k6_tmap_1(X16,X17))=k5_tmap_1(X16,X17)|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t104_tmap_1])])])])])).
cnf(c_0_17, plain, (X1=k6_tmap_1(X2,X4)|v2_struct_0(X2)|X1!=k8_tmap_1(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|X4!=u1_struct_0(X3)|~v1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, plain, (v2_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (l1_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (v1_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_22, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>(v1_funct_1(k2_tmap_1(X1,k8_tmap_1(X1,X2),k9_tmap_1(X1,X2),X2))&(v1_funct_2(k2_tmap_1(X1,k8_tmap_1(X1,X2),k9_tmap_1(X1,X2),X2),u1_struct_0(X2),u1_struct_0(k8_tmap_1(X1,X2)))&(v5_pre_topc(k2_tmap_1(X1,k8_tmap_1(X1,X2),k9_tmap_1(X1,X2),X2),X2,k8_tmap_1(X1,X2))&m1_subset_1(k2_tmap_1(X1,k8_tmap_1(X1,X2),k9_tmap_1(X1,X2),X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k8_tmap_1(X1,X2))))))))))), inference(assume_negation,[status(cth)],[t121_tmap_1])).
fof(c_0_23, plain, ![X20]:(~v2_struct_0(X20)|~l1_struct_0(X20)|v1_xboole_0(u1_struct_0(X20))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_struct_0])])).
fof(c_0_24, plain, ![X21]:(~l1_pre_topc(X21)|l1_struct_0(X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_25, plain, (u1_struct_0(k6_tmap_1(X1,X2))=u1_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26, plain, (k6_tmap_1(X1,u1_struct_0(X2))=k8_tmap_1(X1,X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_17])]), c_0_18]), c_0_19]), c_0_20]), c_0_21])).
fof(c_0_27, negated_conjecture, ((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk2_0))&(~v1_funct_1(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k9_tmap_1(esk2_0,esk3_0),esk3_0))|(~v1_funct_2(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k9_tmap_1(esk2_0,esk3_0),esk3_0),u1_struct_0(esk3_0),u1_struct_0(k8_tmap_1(esk2_0,esk3_0)))|(~v5_pre_topc(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k9_tmap_1(esk2_0,esk3_0),esk3_0),esk3_0,k8_tmap_1(esk2_0,esk3_0))|~m1_subset_1(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k9_tmap_1(esk2_0,esk3_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(k8_tmap_1(esk2_0,esk3_0)))))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_22])])])])).
fof(c_0_28, plain, ![X22]:(v2_struct_0(X22)|~l1_struct_0(X22)|~v1_xboole_0(u1_struct_0(X22))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_29, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, plain, (u1_struct_0(k8_tmap_1(X1,X2))=u1_struct_0(X1)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_21])).
cnf(c_0_32, negated_conjecture, (m1_pre_topc(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_34, negated_conjecture, (v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, plain, (v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_struct_0(X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_38, negated_conjecture, (u1_struct_0(k8_tmap_1(esk2_0,esk3_0))=u1_struct_0(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_39, negated_conjecture, (l1_pre_topc(k8_tmap_1(esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_32]), c_0_33]), c_0_34])]), c_0_35])).
fof(c_0_40, plain, ![X23, X24]:(v2_struct_0(X23)|(~v2_pre_topc(X23)|~l1_pre_topc(X23))|(~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X23)))|k7_tmap_1(X23,X24)=k6_partfun1(u1_struct_0(X23)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_tmap_1])])])])).
cnf(c_0_41, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_36, c_0_30])).
cnf(c_0_42, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|~v2_struct_0(k8_tmap_1(esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])])).
fof(c_0_43, plain, ![X25, X26]:((v1_funct_1(k7_tmap_1(X25,X26))|(v2_struct_0(X25)|(~v2_pre_topc(X25)|~l1_pre_topc(X25))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))))&((v1_funct_2(k7_tmap_1(X25,X26),u1_struct_0(X25),u1_struct_0(k6_tmap_1(X25,X26)))|(v2_struct_0(X25)|(~v2_pre_topc(X25)|~l1_pre_topc(X25))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))))&(m1_subset_1(k7_tmap_1(X25,X26),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X25),u1_struct_0(k6_tmap_1(X25,X26)))))|(v2_struct_0(X25)|(~v2_pre_topc(X25)|~l1_pre_topc(X25))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_tmap_1])])])])).
cnf(c_0_44, plain, (v2_struct_0(X1)|k7_tmap_1(X1,X2)=k6_partfun1(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_45, negated_conjecture, (v2_pre_topc(k8_tmap_1(esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_32]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_46, negated_conjecture, (~v2_struct_0(k8_tmap_1(esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_33])]), c_0_35])).
cnf(c_0_47, plain, (m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_48, negated_conjecture, (k7_tmap_1(k8_tmap_1(esk2_0,esk3_0),X1)=k6_partfun1(u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_38]), c_0_39]), c_0_45])]), c_0_46])).
cnf(c_0_49, plain, (v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_50, negated_conjecture, (m1_subset_1(k6_partfun1(u1_struct_0(esk2_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(k6_tmap_1(k8_tmap_1(esk2_0,esk3_0),X1)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_38]), c_0_38]), c_0_39]), c_0_45])]), c_0_46])).
cnf(c_0_51, negated_conjecture, (v1_funct_2(k6_partfun1(u1_struct_0(esk2_0)),u1_struct_0(esk2_0),u1_struct_0(k6_tmap_1(k8_tmap_1(esk2_0,esk3_0),X1)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_48]), c_0_38]), c_0_38]), c_0_39]), c_0_45])]), c_0_46])).
cnf(c_0_52, plain, (v1_funct_1(k7_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_53, plain, ![X27, X28, X29, X30]:((X29!=k9_tmap_1(X27,X28)|(~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X27)))|(X30!=u1_struct_0(X28)|r1_funct_2(u1_struct_0(X27),u1_struct_0(k8_tmap_1(X27,X28)),u1_struct_0(X27),u1_struct_0(k6_tmap_1(X27,X30)),X29,k7_tmap_1(X27,X30))))|(~v1_funct_1(X29)|(~v1_funct_2(X29,u1_struct_0(X27),u1_struct_0(k8_tmap_1(X27,X28)))|~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X27),u1_struct_0(k8_tmap_1(X27,X28)))))))|~m1_pre_topc(X28,X27)|(v2_struct_0(X27)|(~v2_pre_topc(X27)|~l1_pre_topc(X27))))&((m1_subset_1(esk4_3(X27,X28,X29),k1_zfmisc_1(u1_struct_0(X27)))|X29=k9_tmap_1(X27,X28)|(~v1_funct_1(X29)|(~v1_funct_2(X29,u1_struct_0(X27),u1_struct_0(k8_tmap_1(X27,X28)))|~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X27),u1_struct_0(k8_tmap_1(X27,X28)))))))|~m1_pre_topc(X28,X27)|(v2_struct_0(X27)|(~v2_pre_topc(X27)|~l1_pre_topc(X27))))&((esk4_3(X27,X28,X29)=u1_struct_0(X28)|X29=k9_tmap_1(X27,X28)|(~v1_funct_1(X29)|(~v1_funct_2(X29,u1_struct_0(X27),u1_struct_0(k8_tmap_1(X27,X28)))|~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X27),u1_struct_0(k8_tmap_1(X27,X28)))))))|~m1_pre_topc(X28,X27)|(v2_struct_0(X27)|(~v2_pre_topc(X27)|~l1_pre_topc(X27))))&(~r1_funct_2(u1_struct_0(X27),u1_struct_0(k8_tmap_1(X27,X28)),u1_struct_0(X27),u1_struct_0(k6_tmap_1(X27,esk4_3(X27,X28,X29))),X29,k7_tmap_1(X27,esk4_3(X27,X28,X29)))|X29=k9_tmap_1(X27,X28)|(~v1_funct_1(X29)|(~v1_funct_2(X29,u1_struct_0(X27),u1_struct_0(k8_tmap_1(X27,X28)))|~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X27),u1_struct_0(k8_tmap_1(X27,X28)))))))|~m1_pre_topc(X28,X27)|(v2_struct_0(X27)|(~v2_pre_topc(X27)|~l1_pre_topc(X27))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d12_tmap_1])])])])])])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(k6_partfun1(u1_struct_0(esk2_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_25]), c_0_38]), c_0_38]), c_0_39]), c_0_45])]), c_0_46])).
cnf(c_0_55, negated_conjecture, (v1_funct_2(k6_partfun1(u1_struct_0(esk2_0)),u1_struct_0(esk2_0),u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_25]), c_0_38]), c_0_38]), c_0_39]), c_0_45])]), c_0_46])).
cnf(c_0_56, plain, (v1_funct_1(k6_partfun1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_52, c_0_44])).
cnf(c_0_57, plain, (esk4_3(X1,X2,X3)=u1_struct_0(X2)|X3=k9_tmap_1(X1,X2)|v2_struct_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_58, negated_conjecture, (m1_subset_1(k6_partfun1(u1_struct_0(esk2_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_21]), c_0_33])])).
cnf(c_0_59, negated_conjecture, (v1_funct_2(k6_partfun1(u1_struct_0(esk2_0)),u1_struct_0(esk2_0),u1_struct_0(esk2_0))|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_21]), c_0_33])])).
cnf(c_0_60, plain, (v1_funct_1(k6_partfun1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_56, c_0_21])).
cnf(c_0_61, negated_conjecture, (k7_tmap_1(esk2_0,X1)=k7_tmap_1(k8_tmap_1(esk2_0,esk3_0),X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_48]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_62, negated_conjecture, (esk4_3(esk2_0,esk3_0,X1)=u1_struct_0(esk3_0)|X1=k9_tmap_1(esk2_0,esk3_0)|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk2_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_38]), c_0_32]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_63, negated_conjecture, (m1_subset_1(k6_partfun1(u1_struct_0(esk2_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))), inference(spm,[status(thm)],[c_0_58, c_0_32])).
cnf(c_0_64, negated_conjecture, (v1_funct_2(k6_partfun1(u1_struct_0(esk2_0)),u1_struct_0(esk2_0),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_59, c_0_32])).
cnf(c_0_65, negated_conjecture, (v1_funct_1(k6_partfun1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_32]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_66, negated_conjecture, (k7_tmap_1(esk2_0,X1)=k6_partfun1(u1_struct_0(esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_61]), c_0_38]), c_0_38]), c_0_39]), c_0_45])]), c_0_46])).
cnf(c_0_67, plain, (X3=k9_tmap_1(X1,X2)|v2_struct_0(X1)|~r1_funct_2(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,esk4_3(X1,X2,X3))),X3,k7_tmap_1(X1,esk4_3(X1,X2,X3)))|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_68, negated_conjecture, (esk4_3(esk2_0,esk3_0,k6_partfun1(u1_struct_0(esk2_0)))=u1_struct_0(esk3_0)|k6_partfun1(u1_struct_0(esk2_0))=k9_tmap_1(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64]), c_0_65])])).
cnf(c_0_69, negated_conjecture, (k7_tmap_1(esk2_0,X1)=k6_partfun1(u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_pre_topc(X2,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_21]), c_0_33])])).
cnf(c_0_70, negated_conjecture, (k6_partfun1(u1_struct_0(esk2_0))=k9_tmap_1(esk2_0,esk3_0)|~r1_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk2_0),u1_struct_0(esk2_0),u1_struct_0(k6_tmap_1(esk2_0,u1_struct_0(esk3_0))),k6_partfun1(u1_struct_0(esk2_0)),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_38]), c_0_38]), c_0_64]), c_0_65]), c_0_38]), c_0_63]), c_0_32]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_71, negated_conjecture, (k7_tmap_1(esk2_0,X1)=k6_partfun1(u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_69, c_0_32])).
cnf(c_0_72, negated_conjecture, (k6_partfun1(u1_struct_0(esk2_0))=k9_tmap_1(esk2_0,esk3_0)|~r1_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk2_0),u1_struct_0(esk2_0),u1_struct_0(esk2_0),k6_partfun1(u1_struct_0(esk2_0)),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_26]), c_0_38]), c_0_32]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_73, negated_conjecture, (k7_tmap_1(esk2_0,X1)=k7_tmap_1(esk2_0,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_71]), c_0_33]), c_0_34])]), c_0_35])).
fof(c_0_74, plain, ![X32, X33, X34, X35, X36, X37]:((~r1_funct_2(X32,X33,X34,X35,X36,X37)|X36=X37|(v1_xboole_0(X33)|(v1_xboole_0(X35)|(~v1_funct_1(X36)|(~v1_funct_2(X36,X32,X33)|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))|(~v1_funct_1(X37)|(~v1_funct_2(X37,X34,X35)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(X34,X35)))))))))&(X36!=X37|r1_funct_2(X32,X33,X34,X35,X36,X37)|(v1_xboole_0(X33)|(v1_xboole_0(X35)|(~v1_funct_1(X36)|(~v1_funct_2(X36,X32,X33)|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))|(~v1_funct_1(X37)|(~v1_funct_2(X37,X34,X35)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(X34,X35)))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r1_funct_2])])])])).
cnf(c_0_75, negated_conjecture, (k6_partfun1(u1_struct_0(esk2_0))=k9_tmap_1(esk2_0,esk3_0)|~r1_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk2_0),u1_struct_0(esk2_0),u1_struct_0(esk2_0),k6_partfun1(u1_struct_0(esk2_0)),k7_tmap_1(esk2_0,X1))|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_76, plain, (r1_funct_2(X3,X4,X5,X6,X1,X2)|v1_xboole_0(X4)|v1_xboole_0(X6)|X1!=X2|~v1_funct_1(X1)|~v1_funct_2(X1,X3,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X2)|~v1_funct_2(X2,X5,X6)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_77, negated_conjecture, (k6_partfun1(u1_struct_0(esk2_0))=k9_tmap_1(esk2_0,esk3_0)|~r1_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk2_0),u1_struct_0(esk2_0),u1_struct_0(esk2_0),k6_partfun1(u1_struct_0(esk2_0)),k6_partfun1(u1_struct_0(esk2_0)))|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_44]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_78, plain, (r1_funct_2(X1,X2,X3,X4,X5,X5)|v1_xboole_0(X4)|v1_xboole_0(X2)|~v1_funct_2(X5,X3,X4)|~v1_funct_2(X5,X1,X2)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(er,[status(thm)],[c_0_76])).
cnf(c_0_79, negated_conjecture, (k6_partfun1(u1_struct_0(esk2_0))=k9_tmap_1(esk2_0,esk3_0)|v1_xboole_0(u1_struct_0(esk2_0))|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_64]), c_0_65]), c_0_63])])).
cnf(c_0_80, negated_conjecture, (k6_partfun1(u1_struct_0(esk2_0))=k9_tmap_1(esk2_0,esk3_0)|v1_xboole_0(u1_struct_0(esk2_0))|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_21]), c_0_33])])).
cnf(c_0_81, negated_conjecture, (k6_partfun1(u1_struct_0(esk2_0))=k9_tmap_1(esk2_0,esk3_0)|v1_xboole_0(u1_struct_0(esk2_0))|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_80, c_0_32])).
cnf(c_0_82, negated_conjecture, (k7_tmap_1(esk2_0,X1)=k9_tmap_1(esk2_0,esk3_0)|v1_xboole_0(u1_struct_0(esk2_0))|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_81]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_83, negated_conjecture, (k7_tmap_1(esk2_0,X1)=k9_tmap_1(esk2_0,esk3_0)|v1_xboole_0(u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_21]), c_0_32]), c_0_33])])).
cnf(c_0_84, negated_conjecture, (k7_tmap_1(esk2_0,u1_struct_0(X1))=k9_tmap_1(esk2_0,esk3_0)|v1_xboole_0(u1_struct_0(esk2_0))|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_21]), c_0_33])])).
cnf(c_0_85, negated_conjecture, (k9_tmap_1(esk2_0,esk3_0)=k7_tmap_1(esk2_0,u1_struct_0(esk3_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_84, c_0_32])).
cnf(c_0_86, plain, (m1_subset_1(esk4_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|X3=k9_tmap_1(X1,X2)|v2_struct_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_87, negated_conjecture, (k9_tmap_1(esk2_0,esk3_0)=k7_tmap_1(esk2_0,u1_struct_0(esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_85]), c_0_33])]), c_0_35])).
fof(c_0_88, plain, ![X38, X39, X40]:((v1_funct_1(k2_tmap_1(X38,k6_tmap_1(X38,X39),k7_tmap_1(X38,X39),X40))|u1_struct_0(X40)!=X39|(v2_struct_0(X40)|~m1_pre_topc(X40,X38))|~m1_subset_1(X39,k1_zfmisc_1(u1_struct_0(X38)))|(v2_struct_0(X38)|(~v2_pre_topc(X38)|~l1_pre_topc(X38))))&((v1_funct_2(k2_tmap_1(X38,k6_tmap_1(X38,X39),k7_tmap_1(X38,X39),X40),u1_struct_0(X40),u1_struct_0(k6_tmap_1(X38,X39)))|u1_struct_0(X40)!=X39|(v2_struct_0(X40)|~m1_pre_topc(X40,X38))|~m1_subset_1(X39,k1_zfmisc_1(u1_struct_0(X38)))|(v2_struct_0(X38)|(~v2_pre_topc(X38)|~l1_pre_topc(X38))))&((v5_pre_topc(k2_tmap_1(X38,k6_tmap_1(X38,X39),k7_tmap_1(X38,X39),X40),X40,k6_tmap_1(X38,X39))|u1_struct_0(X40)!=X39|(v2_struct_0(X40)|~m1_pre_topc(X40,X38))|~m1_subset_1(X39,k1_zfmisc_1(u1_struct_0(X38)))|(v2_struct_0(X38)|(~v2_pre_topc(X38)|~l1_pre_topc(X38))))&(m1_subset_1(k2_tmap_1(X38,k6_tmap_1(X38,X39),k7_tmap_1(X38,X39),X40),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X40),u1_struct_0(k6_tmap_1(X38,X39)))))|u1_struct_0(X40)!=X39|(v2_struct_0(X40)|~m1_pre_topc(X40,X38))|~m1_subset_1(X39,k1_zfmisc_1(u1_struct_0(X38)))|(v2_struct_0(X38)|(~v2_pre_topc(X38)|~l1_pre_topc(X38))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t112_tmap_1])])])])])).
cnf(c_0_89, negated_conjecture, (X1=k7_tmap_1(esk2_0,u1_struct_0(esk3_0))|m1_subset_1(esk4_3(esk2_0,esk3_0,X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk2_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_38]), c_0_32]), c_0_33]), c_0_34])]), c_0_35]), c_0_87])).
cnf(c_0_90, negated_conjecture, (esk4_3(esk2_0,esk3_0,k6_partfun1(u1_struct_0(esk2_0)))=u1_struct_0(esk3_0)|k6_partfun1(u1_struct_0(esk2_0))=k7_tmap_1(esk2_0,u1_struct_0(esk3_0))), inference(rw,[status(thm)],[c_0_68, c_0_87])).
cnf(c_0_91, plain, (v1_funct_1(k2_tmap_1(X1,k6_tmap_1(X1,X2),k7_tmap_1(X1,X2),X3))|v2_struct_0(X3)|v2_struct_0(X1)|u1_struct_0(X3)!=X2|~m1_pre_topc(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_92, negated_conjecture, (k6_partfun1(u1_struct_0(esk2_0))=k7_tmap_1(esk2_0,u1_struct_0(esk3_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_64]), c_0_65]), c_0_63])]), c_0_71])).
cnf(c_0_93, plain, (v1_funct_1(k2_tmap_1(X1,k6_tmap_1(X1,u1_struct_0(X2)),k7_tmap_1(X1,u1_struct_0(X2)),X2))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_91]), c_0_21])).
cnf(c_0_94, negated_conjecture, (k7_tmap_1(esk2_0,X1)=k7_tmap_1(esk2_0,u1_struct_0(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(rw,[status(thm)],[c_0_71, c_0_92])).
cnf(c_0_95, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_96, plain, (v5_pre_topc(k2_tmap_1(X1,k6_tmap_1(X1,X2),k7_tmap_1(X1,X2),X3),X3,k6_tmap_1(X1,X2))|v2_struct_0(X3)|v2_struct_0(X1)|u1_struct_0(X3)!=X2|~m1_pre_topc(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_97, negated_conjecture, (v1_funct_1(k2_tmap_1(esk2_0,k6_tmap_1(esk2_0,u1_struct_0(esk3_0)),k7_tmap_1(esk2_0,X1),esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_32]), c_0_33]), c_0_34])]), c_0_35]), c_0_95])).
cnf(c_0_98, plain, (v5_pre_topc(k2_tmap_1(X1,k6_tmap_1(X1,u1_struct_0(X2)),k7_tmap_1(X1,u1_struct_0(X2)),X2),X2,k6_tmap_1(X1,u1_struct_0(X2)))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_96]), c_0_21])).
cnf(c_0_99, plain, (v1_funct_2(k2_tmap_1(X1,k6_tmap_1(X1,X2),k7_tmap_1(X1,X2),X3),u1_struct_0(X3),u1_struct_0(k6_tmap_1(X1,X2)))|v2_struct_0(X3)|v2_struct_0(X1)|u1_struct_0(X3)!=X2|~m1_pre_topc(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_100, negated_conjecture, (v1_funct_1(k2_tmap_1(esk2_0,k6_tmap_1(esk2_0,u1_struct_0(esk3_0)),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_44]), c_0_92]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_101, negated_conjecture, (v5_pre_topc(k2_tmap_1(esk2_0,k6_tmap_1(esk2_0,u1_struct_0(esk3_0)),k7_tmap_1(esk2_0,X1),esk3_0),esk3_0,k6_tmap_1(esk2_0,u1_struct_0(esk3_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_94]), c_0_32]), c_0_33]), c_0_34])]), c_0_35]), c_0_95])).
cnf(c_0_102, plain, (v1_funct_2(k2_tmap_1(X1,k6_tmap_1(X1,u1_struct_0(X2)),k7_tmap_1(X1,u1_struct_0(X2)),X2),u1_struct_0(X2),u1_struct_0(k6_tmap_1(X1,u1_struct_0(X2))))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_99]), c_0_21])).
cnf(c_0_103, negated_conjecture, (~v1_funct_1(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k9_tmap_1(esk2_0,esk3_0),esk3_0))|~v1_funct_2(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k9_tmap_1(esk2_0,esk3_0),esk3_0),u1_struct_0(esk3_0),u1_struct_0(k8_tmap_1(esk2_0,esk3_0)))|~v5_pre_topc(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k9_tmap_1(esk2_0,esk3_0),esk3_0),esk3_0,k8_tmap_1(esk2_0,esk3_0))|~m1_subset_1(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k9_tmap_1(esk2_0,esk3_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(k8_tmap_1(esk2_0,esk3_0)))))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_104, negated_conjecture, (v1_funct_1(k2_tmap_1(esk2_0,k6_tmap_1(esk2_0,u1_struct_0(esk3_0)),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0))|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_21]), c_0_33])])).
cnf(c_0_105, negated_conjecture, (v5_pre_topc(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k7_tmap_1(esk2_0,X1),esk3_0),esk3_0,k8_tmap_1(esk2_0,esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_26]), c_0_32]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_106, negated_conjecture, (v1_funct_2(k2_tmap_1(esk2_0,k6_tmap_1(esk2_0,u1_struct_0(esk3_0)),k7_tmap_1(esk2_0,X1),esk3_0),u1_struct_0(esk3_0),u1_struct_0(k6_tmap_1(esk2_0,u1_struct_0(esk3_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_94]), c_0_32]), c_0_33]), c_0_34])]), c_0_35]), c_0_95])).
cnf(c_0_107, negated_conjecture, (~v5_pre_topc(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k9_tmap_1(esk2_0,esk3_0),esk3_0),esk3_0,k8_tmap_1(esk2_0,esk3_0))|~v1_funct_2(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k9_tmap_1(esk2_0,esk3_0),esk3_0),u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_1(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k9_tmap_1(esk2_0,esk3_0),esk3_0))|~m1_subset_1(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k9_tmap_1(esk2_0,esk3_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_38]), c_0_38])).
cnf(c_0_108, negated_conjecture, (v1_funct_1(k2_tmap_1(esk2_0,k6_tmap_1(esk2_0,u1_struct_0(esk3_0)),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0))), inference(spm,[status(thm)],[c_0_104, c_0_32])).
cnf(c_0_109, negated_conjecture, (v5_pre_topc(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0),esk3_0,k8_tmap_1(esk2_0,esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_44]), c_0_92]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_110, negated_conjecture, (v1_funct_2(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k7_tmap_1(esk2_0,X1),esk3_0),u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_26]), c_0_38]), c_0_32]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_111, negated_conjecture, (~v5_pre_topc(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0),esk3_0,k8_tmap_1(esk2_0,esk3_0))|~v1_funct_2(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0),u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_1(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0))|~m1_subset_1(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_107, c_0_87]), c_0_87]), c_0_87]), c_0_87])).
cnf(c_0_112, negated_conjecture, (v1_funct_1(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_26]), c_0_32]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_113, negated_conjecture, (v5_pre_topc(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0),esk3_0,k8_tmap_1(esk2_0,esk3_0))|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_21]), c_0_33])])).
cnf(c_0_114, negated_conjecture, (v1_funct_2(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0),u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_44]), c_0_92]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_115, plain, (m1_subset_1(k2_tmap_1(X1,k6_tmap_1(X1,X2),k7_tmap_1(X1,X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k6_tmap_1(X1,X2)))))|v2_struct_0(X3)|v2_struct_0(X1)|u1_struct_0(X3)!=X2|~m1_pre_topc(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_116, negated_conjecture, (~v5_pre_topc(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0),esk3_0,k8_tmap_1(esk2_0,esk3_0))|~v1_funct_2(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0),u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~m1_subset_1(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_111, c_0_112])])).
cnf(c_0_117, negated_conjecture, (v5_pre_topc(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0),esk3_0,k8_tmap_1(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_113, c_0_32])).
cnf(c_0_118, negated_conjecture, (v1_funct_2(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0),u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_21]), c_0_33])])).
cnf(c_0_119, plain, (m1_subset_1(k2_tmap_1(X1,k6_tmap_1(X1,u1_struct_0(X2)),k7_tmap_1(X1,u1_struct_0(X2)),X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k6_tmap_1(X1,u1_struct_0(X2))))))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_115]), c_0_21])).
cnf(c_0_120, negated_conjecture, (~v1_funct_2(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0),u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~m1_subset_1(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_116, c_0_117])])).
cnf(c_0_121, negated_conjecture, (v1_funct_2(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0),u1_struct_0(esk3_0),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_118, c_0_32])).
cnf(c_0_122, plain, (m1_subset_1(k2_tmap_1(X1,k8_tmap_1(X1,X2),k7_tmap_1(X1,u1_struct_0(X2)),X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k8_tmap_1(X1,X2)))))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_119, c_0_26])).
cnf(c_0_123, negated_conjecture, (~m1_subset_1(k2_tmap_1(esk2_0,k8_tmap_1(esk2_0,esk3_0),k7_tmap_1(esk2_0,u1_struct_0(esk3_0)),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_120, c_0_121])])).
cnf(c_0_124, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_38]), c_0_32]), c_0_33]), c_0_34])]), c_0_123]), c_0_95]), c_0_35]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 125
# Proof object clause steps            : 98
# Proof object formula steps           : 27
# Proof object conjectures             : 68
# Proof object clause conjectures      : 65
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 58
# Proof object simplifying inferences  : 219
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 33
# Processed clauses                    : 1150
# ...of these trivial                  : 5
# ...subsumed                          : 782
# ...remaining for further processing  : 363
# Other redundant clauses eliminated   : 10
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 72
# Backward-rewritten                   : 54
# Generated clauses                    : 2122
# ...of the previous two non-trivial   : 2043
# Contextual simplify-reflections      : 25
# Paramodulations                      : 2114
# Factorizations                       : 0
# Equation resolutions                 : 10
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 197
#    Positive orientable unit clauses  : 18
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 175
# Current number of unprocessed clauses: 814
# ...number of literals in the above   : 5847
# Current number of archived formulas  : 0
# Current number of archived clauses   : 159
# Clause-clause subsumption calls (NU) : 23902
# Rec. Clause-clause subsumption calls : 8317
# Non-unit clause-clause subsumptions  : 873
# Unit Clause-clause subsumption calls : 496
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 31
# BW rewrite match successes           : 12
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 77584

# -------------------------------------------------
# User time                : 0.048 s
# System time              : 0.007 s
# Total time               : 0.055 s
# Maximum resident set size: 28696 pages
