# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t50_lattice3, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))&k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0))), file('data/MPTP/proofs/0a47d3fca72e479d82202f73a5fb45cc.E_input', t50_lattice3)).
fof(t48_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X2)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>~((r3_lattices(X1,X4,X5)&r2_hidden(X5,X3)))))))=>r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3)))), file('data/MPTP/proofs/0a47d3fca72e479d82202f73a5fb45cc.E_input', t48_lattice3)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/0a47d3fca72e479d82202f73a5fb45cc.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/0a47d3fca72e479d82202f73a5fb45cc.E_input', t2_xboole_1)).
fof(t21_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattices(X1,X2,X3)<=>k2_lattices(X1,X2,X3)=X2)))), file('data/MPTP/proofs/0a47d3fca72e479d82202f73a5fb45cc.E_input', t21_lattices)).
fof(redefinition_r3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_lattices(X1,X2,X3)<=>r1_lattices(X1,X2,X3))), file('data/MPTP/proofs/0a47d3fca72e479d82202f73a5fb45cc.E_input', redefinition_r3_lattices)).
fof(t45_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(X2=k15_lattice3(X1,a_2_3_lattice3(X1,X2))&X2=k16_lattice3(X1,a_2_4_lattice3(X1,X2))))), file('data/MPTP/proofs/0a47d3fca72e479d82202f73a5fb45cc.E_input', t45_lattice3)).
fof(d13_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v13_lattices(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_lattices(X1,X2,X3)=X2&k2_lattices(X1,X3,X2)=X2))))), file('data/MPTP/proofs/0a47d3fca72e479d82202f73a5fb45cc.E_input', d13_lattices)).
fof(d6_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v6_lattices(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2))))), file('data/MPTP/proofs/0a47d3fca72e479d82202f73a5fb45cc.E_input', d6_lattices)).
fof(dt_k15_lattice3, axiom, ![X1, X2]:((~(v2_struct_0(X1))&l3_lattices(X1))=>m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/0a47d3fca72e479d82202f73a5fb45cc.E_input', dt_k15_lattice3)).
fof(d16_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v13_lattices(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(X2=k5_lattices(X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_lattices(X1,X2,X3)=X2&k2_lattices(X1,X3,X2)=X2)))))), file('data/MPTP/proofs/0a47d3fca72e479d82202f73a5fb45cc.E_input', d16_lattices)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/0a47d3fca72e479d82202f73a5fb45cc.E_input', cc1_lattices)).
fof(dt_l3_lattices, axiom, ![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1))), file('data/MPTP/proofs/0a47d3fca72e479d82202f73a5fb45cc.E_input', dt_l3_lattices)).
fof(c_0_13, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))&k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0)))), inference(assume_negation,[status(cth)],[t50_lattice3])).
fof(c_0_14, plain, ![X7, X8, X9, X11]:((m1_subset_1(esk2_3(X7,X8,X9),u1_struct_0(X7))|r3_lattices(X7,k15_lattice3(X7,X8),k15_lattice3(X7,X9))|(v2_struct_0(X7)|(~v10_lattices(X7)|(~v4_lattice3(X7)|~l3_lattices(X7)))))&((r2_hidden(esk2_3(X7,X8,X9),X8)|r3_lattices(X7,k15_lattice3(X7,X8),k15_lattice3(X7,X9))|(v2_struct_0(X7)|(~v10_lattices(X7)|(~v4_lattice3(X7)|~l3_lattices(X7)))))&(~m1_subset_1(X11,u1_struct_0(X7))|(~r3_lattices(X7,esk2_3(X7,X8,X9),X11)|~r2_hidden(X11,X9))|r3_lattices(X7,k15_lattice3(X7,X8),k15_lattice3(X7,X9))|(v2_struct_0(X7)|(~v10_lattices(X7)|(~v4_lattice3(X7)|~l3_lattices(X7))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t48_lattice3])])])])])])).
fof(c_0_15, negated_conjecture, ((~v2_struct_0(esk1_0)&(v10_lattices(esk1_0)&(v4_lattice3(esk1_0)&l3_lattices(esk1_0))))&(v2_struct_0(esk1_0)|(~v10_lattices(esk1_0)|(~v13_lattices(esk1_0)|~l3_lattices(esk1_0)))|k5_lattices(esk1_0)!=k15_lattice3(esk1_0,k1_xboole_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_13])])])])).
fof(c_0_16, plain, ![X12, X13]:(~r2_hidden(X12,X13)|~r1_tarski(X13,X12)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_17, plain, (r2_hidden(esk2_3(X1,X2,X3),X2)|r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3))|v2_struct_0(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, negated_conjecture, (l3_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, negated_conjecture, (v4_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, negated_conjecture, (v10_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, negated_conjecture, (r3_lattices(esk1_0,k15_lattice3(esk1_0,X1),k15_lattice3(esk1_0,X2))|r2_hidden(esk2_3(esk1_0,X1,X2),X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19]), c_0_20])]), c_0_21])).
fof(c_0_24, plain, ![X14]:r1_tarski(k1_xboole_0,X14), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_25, plain, ![X15, X16, X17]:((~r1_lattices(X15,X16,X17)|k2_lattices(X15,X16,X17)=X16|~m1_subset_1(X17,u1_struct_0(X15))|~m1_subset_1(X16,u1_struct_0(X15))|(v2_struct_0(X15)|(~v8_lattices(X15)|(~v9_lattices(X15)|~l3_lattices(X15)))))&(k2_lattices(X15,X16,X17)!=X16|r1_lattices(X15,X16,X17)|~m1_subset_1(X17,u1_struct_0(X15))|~m1_subset_1(X16,u1_struct_0(X15))|(v2_struct_0(X15)|(~v8_lattices(X15)|(~v9_lattices(X15)|~l3_lattices(X15)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t21_lattices])])])])])).
fof(c_0_26, plain, ![X18, X19, X20]:((~r3_lattices(X18,X19,X20)|r1_lattices(X18,X19,X20)|(v2_struct_0(X18)|(~v6_lattices(X18)|(~v8_lattices(X18)|(~v9_lattices(X18)|~l3_lattices(X18))))|(~m1_subset_1(X19,u1_struct_0(X18))|~m1_subset_1(X20,u1_struct_0(X18)))))&(~r1_lattices(X18,X19,X20)|r3_lattices(X18,X19,X20)|(v2_struct_0(X18)|(~v6_lattices(X18)|(~v8_lattices(X18)|(~v9_lattices(X18)|~l3_lattices(X18))))|(~m1_subset_1(X19,u1_struct_0(X18))|~m1_subset_1(X20,u1_struct_0(X18)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_lattices])])])])).
cnf(c_0_27, negated_conjecture, (r3_lattices(esk1_0,k15_lattice3(esk1_0,X1),k15_lattice3(esk1_0,X2))|~r1_tarski(X1,esk2_3(esk1_0,X1,X2))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_28, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_29, plain, ![X21, X22]:((X22=k15_lattice3(X21,a_2_3_lattice3(X21,X22))|~m1_subset_1(X22,u1_struct_0(X21))|(v2_struct_0(X21)|(~v10_lattices(X21)|(~v4_lattice3(X21)|~l3_lattices(X21)))))&(X22=k16_lattice3(X21,a_2_4_lattice3(X21,X22))|~m1_subset_1(X22,u1_struct_0(X21))|(v2_struct_0(X21)|(~v10_lattices(X21)|(~v4_lattice3(X21)|~l3_lattices(X21)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t45_lattice3])])])])])).
fof(c_0_30, plain, ![X23, X25, X26]:(((m1_subset_1(esk3_1(X23),u1_struct_0(X23))|~v13_lattices(X23)|(v2_struct_0(X23)|~l1_lattices(X23)))&((k2_lattices(X23,esk3_1(X23),X25)=esk3_1(X23)|~m1_subset_1(X25,u1_struct_0(X23))|~v13_lattices(X23)|(v2_struct_0(X23)|~l1_lattices(X23)))&(k2_lattices(X23,X25,esk3_1(X23))=esk3_1(X23)|~m1_subset_1(X25,u1_struct_0(X23))|~v13_lattices(X23)|(v2_struct_0(X23)|~l1_lattices(X23)))))&((m1_subset_1(esk4_2(X23,X26),u1_struct_0(X23))|~m1_subset_1(X26,u1_struct_0(X23))|v13_lattices(X23)|(v2_struct_0(X23)|~l1_lattices(X23)))&(k2_lattices(X23,X26,esk4_2(X23,X26))!=X26|k2_lattices(X23,esk4_2(X23,X26),X26)!=X26|~m1_subset_1(X26,u1_struct_0(X23))|v13_lattices(X23)|(v2_struct_0(X23)|~l1_lattices(X23))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_lattices])])])])])])).
fof(c_0_31, plain, ![X28, X29, X30]:((~v6_lattices(X28)|(~m1_subset_1(X29,u1_struct_0(X28))|(~m1_subset_1(X30,u1_struct_0(X28))|k2_lattices(X28,X29,X30)=k2_lattices(X28,X30,X29)))|(v2_struct_0(X28)|~l1_lattices(X28)))&((m1_subset_1(esk5_1(X28),u1_struct_0(X28))|v6_lattices(X28)|(v2_struct_0(X28)|~l1_lattices(X28)))&((m1_subset_1(esk6_1(X28),u1_struct_0(X28))|v6_lattices(X28)|(v2_struct_0(X28)|~l1_lattices(X28)))&(k2_lattices(X28,esk5_1(X28),esk6_1(X28))!=k2_lattices(X28,esk6_1(X28),esk5_1(X28))|v6_lattices(X28)|(v2_struct_0(X28)|~l1_lattices(X28)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_lattices])])])])])])).
cnf(c_0_32, plain, (k2_lattices(X1,X2,X3)=X2|v2_struct_0(X1)|~r1_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (r1_lattices(X1,X2,X3)|v2_struct_0(X1)|~r3_lattices(X1,X2,X3)|~v6_lattices(X1)|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, negated_conjecture, (r3_lattices(esk1_0,k15_lattice3(esk1_0,k1_xboole_0),k15_lattice3(esk1_0,X1))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_35, plain, (X1=k15_lattice3(X2,a_2_3_lattice3(X2,X1))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, plain, (v13_lattices(X1)|v2_struct_0(X1)|k2_lattices(X1,X2,esk4_2(X1,X2))!=X2|k2_lattices(X1,esk4_2(X1,X2),X2)!=X2|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37, plain, (k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2)|v2_struct_0(X1)|~v6_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, plain, (m1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|v13_lattices(X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39, plain, (k2_lattices(X1,X2,X3)=X2|v2_struct_0(X1)|~v6_lattices(X1)|~v9_lattices(X1)|~v8_lattices(X1)|~r3_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_40, negated_conjecture, (r3_lattices(esk1_0,k15_lattice3(esk1_0,k1_xboole_0),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_18]), c_0_19]), c_0_20])]), c_0_21])).
cnf(c_0_41, plain, (v13_lattices(X1)|v2_struct_0(X1)|k2_lattices(X1,X2,esk4_2(X1,X2))!=X2|~l1_lattices(X1)|~v6_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_42, negated_conjecture, (k2_lattices(esk1_0,k15_lattice3(esk1_0,k1_xboole_0),X1)=k15_lattice3(esk1_0,k1_xboole_0)|~v6_lattices(esk1_0)|~v9_lattices(esk1_0)|~v8_lattices(esk1_0)|~m1_subset_1(k15_lattice3(esk1_0,k1_xboole_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_18])]), c_0_21])).
cnf(c_0_43, negated_conjecture, (v13_lattices(esk1_0)|~l1_lattices(esk1_0)|~v6_lattices(esk1_0)|~v9_lattices(esk1_0)|~v8_lattices(esk1_0)|~m1_subset_1(esk4_2(esk1_0,k15_lattice3(esk1_0,k1_xboole_0)),u1_struct_0(esk1_0))|~m1_subset_1(k15_lattice3(esk1_0,k1_xboole_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_21])).
fof(c_0_44, plain, ![X33, X34]:(v2_struct_0(X33)|~l3_lattices(X33)|m1_subset_1(k15_lattice3(X33,X34),u1_struct_0(X33))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k15_lattice3])])])).
fof(c_0_45, plain, ![X35, X36, X37]:(((k2_lattices(X35,X36,X37)=X36|~m1_subset_1(X37,u1_struct_0(X35))|X36!=k5_lattices(X35)|~m1_subset_1(X36,u1_struct_0(X35))|~v13_lattices(X35)|(v2_struct_0(X35)|~l1_lattices(X35)))&(k2_lattices(X35,X37,X36)=X36|~m1_subset_1(X37,u1_struct_0(X35))|X36!=k5_lattices(X35)|~m1_subset_1(X36,u1_struct_0(X35))|~v13_lattices(X35)|(v2_struct_0(X35)|~l1_lattices(X35))))&((m1_subset_1(esk7_2(X35,X36),u1_struct_0(X35))|X36=k5_lattices(X35)|~m1_subset_1(X36,u1_struct_0(X35))|~v13_lattices(X35)|(v2_struct_0(X35)|~l1_lattices(X35)))&(k2_lattices(X35,X36,esk7_2(X35,X36))!=X36|k2_lattices(X35,esk7_2(X35,X36),X36)!=X36|X36=k5_lattices(X35)|~m1_subset_1(X36,u1_struct_0(X35))|~v13_lattices(X35)|(v2_struct_0(X35)|~l1_lattices(X35))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattices])])])])])])).
cnf(c_0_46, negated_conjecture, (v13_lattices(esk1_0)|~l1_lattices(esk1_0)|~v6_lattices(esk1_0)|~v9_lattices(esk1_0)|~v8_lattices(esk1_0)|~m1_subset_1(k15_lattice3(esk1_0,k1_xboole_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_38]), c_0_21])).
cnf(c_0_47, plain, (v2_struct_0(X1)|m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_48, plain, (X2=k5_lattices(X1)|v2_struct_0(X1)|k2_lattices(X1,X2,esk7_2(X1,X2))!=X2|k2_lattices(X1,esk7_2(X1,X2),X2)!=X2|~m1_subset_1(X2,u1_struct_0(X1))|~v13_lattices(X1)|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_49, negated_conjecture, (k2_lattices(esk1_0,X1,k15_lattice3(esk1_0,k1_xboole_0))=k15_lattice3(esk1_0,k1_xboole_0)|~l1_lattices(esk1_0)|~v6_lattices(esk1_0)|~v9_lattices(esk1_0)|~v8_lattices(esk1_0)|~m1_subset_1(k15_lattice3(esk1_0,k1_xboole_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_42]), c_0_21])).
cnf(c_0_50, negated_conjecture, (v13_lattices(esk1_0)|~l1_lattices(esk1_0)|~v6_lattices(esk1_0)|~v9_lattices(esk1_0)|~v8_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_18])]), c_0_21])).
cnf(c_0_51, negated_conjecture, (k15_lattice3(esk1_0,k1_xboole_0)=k5_lattices(esk1_0)|~l1_lattices(esk1_0)|~v6_lattices(esk1_0)|~v9_lattices(esk1_0)|~v8_lattices(esk1_0)|~m1_subset_1(esk7_2(esk1_0,k15_lattice3(esk1_0,k1_xboole_0)),u1_struct_0(esk1_0))|~m1_subset_1(k15_lattice3(esk1_0,k1_xboole_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_21]), c_0_50]), c_0_42])).
cnf(c_0_52, plain, (m1_subset_1(esk7_2(X1,X2),u1_struct_0(X1))|X2=k5_lattices(X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v13_lattices(X1)|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_53, negated_conjecture, (v2_struct_0(esk1_0)|~v10_lattices(esk1_0)|~v13_lattices(esk1_0)|~l3_lattices(esk1_0)|k5_lattices(esk1_0)!=k15_lattice3(esk1_0,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_54, negated_conjecture, (k15_lattice3(esk1_0,k1_xboole_0)=k5_lattices(esk1_0)|~l1_lattices(esk1_0)|~v6_lattices(esk1_0)|~v9_lattices(esk1_0)|~v8_lattices(esk1_0)|~m1_subset_1(k15_lattice3(esk1_0,k1_xboole_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_21]), c_0_50])).
cnf(c_0_55, negated_conjecture, (k15_lattice3(esk1_0,k1_xboole_0)!=k5_lattices(esk1_0)|~v13_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_20]), c_0_18])]), c_0_21])).
cnf(c_0_56, negated_conjecture, (k15_lattice3(esk1_0,k1_xboole_0)=k5_lattices(esk1_0)|~l1_lattices(esk1_0)|~v6_lattices(esk1_0)|~v9_lattices(esk1_0)|~v8_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_47]), c_0_18])]), c_0_21])).
fof(c_0_57, plain, ![X39]:((~v2_struct_0(X39)|(v2_struct_0(X39)|~v10_lattices(X39))|~l3_lattices(X39))&((v4_lattices(X39)|(v2_struct_0(X39)|~v10_lattices(X39))|~l3_lattices(X39))&((v5_lattices(X39)|(v2_struct_0(X39)|~v10_lattices(X39))|~l3_lattices(X39))&((v6_lattices(X39)|(v2_struct_0(X39)|~v10_lattices(X39))|~l3_lattices(X39))&((v7_lattices(X39)|(v2_struct_0(X39)|~v10_lattices(X39))|~l3_lattices(X39))&((v8_lattices(X39)|(v2_struct_0(X39)|~v10_lattices(X39))|~l3_lattices(X39))&(v9_lattices(X39)|(v2_struct_0(X39)|~v10_lattices(X39))|~l3_lattices(X39)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
cnf(c_0_58, negated_conjecture, (~l1_lattices(esk1_0)|~v6_lattices(esk1_0)|~v9_lattices(esk1_0)|~v8_lattices(esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_50])).
cnf(c_0_59, plain, (v9_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_60, negated_conjecture, (~l1_lattices(esk1_0)|~v6_lattices(esk1_0)|~v8_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_18]), c_0_20])]), c_0_21])).
cnf(c_0_61, plain, (v8_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_62, negated_conjecture, (~l1_lattices(esk1_0)|~v6_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_18]), c_0_20])]), c_0_21])).
cnf(c_0_63, plain, (v6_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
fof(c_0_64, plain, ![X40]:((l1_lattices(X40)|~l3_lattices(X40))&(l2_lattices(X40)|~l3_lattices(X40))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
cnf(c_0_65, negated_conjecture, (~l1_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_18]), c_0_20])]), c_0_21])).
cnf(c_0_66, plain, (l1_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_67, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_18])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 68
# Proof object clause steps            : 41
# Proof object formula steps           : 27
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 46
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 38
# Processed clauses                    : 99
# ...of these trivial                  : 0
# ...subsumed                          : 14
# ...remaining for further processing  : 85
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 25
# Backward-rewritten                   : 0
# Generated clauses                    : 114
# ...of the previous two non-trivial   : 103
# Contextual simplify-reflections      : 11
# Paramodulations                      : 112
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 58
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 50
# Current number of unprocessed clauses: 42
# ...number of literals in the above   : 281
# Current number of archived formulas  : 0
# Current number of archived clauses   : 25
# Clause-clause subsumption calls (NU) : 1872
# Rec. Clause-clause subsumption calls : 183
# Non-unit clause-clause subsumptions  : 49
# Unit Clause-clause subsumption calls : 23
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6825

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 31508 pages
