# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t204_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)<=>r2_hidden(X2,k11_relat_1(X3,X1)))), file('data/MPTP/proofs/0a5c284b5d8c4d18ba01b562df83b1dc.E_input', t204_relat_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/0a5c284b5d8c4d18ba01b562df83b1dc.E_input', t3_xboole_0)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/0a5c284b5d8c4d18ba01b562df83b1dc.E_input', d4_relat_1)).
fof(t205_relat_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>(r2_hidden(X1,k1_relat_1(X2))<=>~(k11_relat_1(X2,X1)=k1_xboole_0))), file('data/MPTP/proofs/0a5c284b5d8c4d18ba01b562df83b1dc.E_input', t205_relat_1)).
fof(t151_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(k9_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k1_relat_1(X2),X1))), file('data/MPTP/proofs/0a5c284b5d8c4d18ba01b562df83b1dc.E_input', t151_relat_1)).
fof(d16_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:k11_relat_1(X1,X2)=k9_relat_1(X1,k1_tarski(X2))), file('data/MPTP/proofs/0a5c284b5d8c4d18ba01b562df83b1dc.E_input', d16_relat_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/0a5c284b5d8c4d18ba01b562df83b1dc.E_input', t69_enumset1)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/0a5c284b5d8c4d18ba01b562df83b1dc.E_input', d2_tarski)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/0a5c284b5d8c4d18ba01b562df83b1dc.E_input', t7_xboole_0)).
fof(c_0_9, plain, ![X10, X11, X12]:((~r2_hidden(k4_tarski(X10,X11),X12)|r2_hidden(X11,k11_relat_1(X12,X10))|~v1_relat_1(X12))&(~r2_hidden(X11,k11_relat_1(X12,X10))|r2_hidden(k4_tarski(X10,X11),X12)|~v1_relat_1(X12))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t204_relat_1])])])).
fof(c_0_10, plain, ![X37, X38, X40, X41, X42]:(((r2_hidden(esk8_2(X37,X38),X37)|r1_xboole_0(X37,X38))&(r2_hidden(esk8_2(X37,X38),X38)|r1_xboole_0(X37,X38)))&(~r2_hidden(X42,X40)|~r2_hidden(X42,X41)|~r1_xboole_0(X40,X41))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_11, plain, ![X22, X23, X24, X26, X27, X28, X29, X31]:(((~r2_hidden(X24,X23)|r2_hidden(k4_tarski(X24,esk4_3(X22,X23,X24)),X22)|X23!=k1_relat_1(X22))&(~r2_hidden(k4_tarski(X26,X27),X22)|r2_hidden(X26,X23)|X23!=k1_relat_1(X22)))&((~r2_hidden(esk5_2(X28,X29),X29)|~r2_hidden(k4_tarski(esk5_2(X28,X29),X31),X28)|X29=k1_relat_1(X28))&(r2_hidden(esk5_2(X28,X29),X29)|r2_hidden(k4_tarski(esk5_2(X28,X29),esk6_2(X28,X29)),X28)|X29=k1_relat_1(X28)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
cnf(c_0_12, plain, (r2_hidden(k4_tarski(X3,X1),X2)|~r2_hidden(X1,k11_relat_1(X2,X3))|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (r2_hidden(esk8_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (r1_xboole_0(X1,k11_relat_1(X2,X3))|r2_hidden(k4_tarski(X3,esk8_2(X1,k11_relat_1(X2,X3))),X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
fof(c_0_16, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>(r2_hidden(X1,k1_relat_1(X2))<=>~(k11_relat_1(X2,X1)=k1_xboole_0)))), inference(assume_negation,[status(cth)],[t205_relat_1])).
cnf(c_0_17, plain, (r1_xboole_0(X1,k11_relat_1(X2,X3))|r2_hidden(X3,X4)|X4!=k1_relat_1(X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
fof(c_0_18, negated_conjecture, (v1_relat_1(esk2_0)&((~r2_hidden(esk1_0,k1_relat_1(esk2_0))|k11_relat_1(esk2_0,esk1_0)=k1_xboole_0)&(r2_hidden(esk1_0,k1_relat_1(esk2_0))|k11_relat_1(esk2_0,esk1_0)!=k1_xboole_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_16])])])])).
fof(c_0_19, plain, ![X8, X9]:((k9_relat_1(X9,X8)!=k1_xboole_0|r1_xboole_0(k1_relat_1(X9),X8)|~v1_relat_1(X9))&(~r1_xboole_0(k1_relat_1(X9),X8)|k9_relat_1(X9,X8)=k1_xboole_0|~v1_relat_1(X9))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t151_relat_1])])])).
fof(c_0_20, plain, ![X33, X34]:(~v1_relat_1(X33)|k11_relat_1(X33,X34)=k9_relat_1(X33,k1_tarski(X34))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d16_relat_1])])])).
fof(c_0_21, plain, ![X7]:k2_tarski(X7,X7)=k1_tarski(X7), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_22, plain, (r1_xboole_0(X1,k11_relat_1(X2,X3))|r2_hidden(X3,k1_relat_1(X2))|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_23, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_24, plain, ![X13, X14, X15, X16, X17, X18, X19, X20]:(((~r2_hidden(X16,X15)|(X16=X13|X16=X14)|X15!=k2_tarski(X13,X14))&((X17!=X13|r2_hidden(X17,X15)|X15!=k2_tarski(X13,X14))&(X17!=X14|r2_hidden(X17,X15)|X15!=k2_tarski(X13,X14))))&(((esk3_3(X18,X19,X20)!=X18|~r2_hidden(esk3_3(X18,X19,X20),X20)|X20=k2_tarski(X18,X19))&(esk3_3(X18,X19,X20)!=X19|~r2_hidden(esk3_3(X18,X19,X20),X20)|X20=k2_tarski(X18,X19)))&(r2_hidden(esk3_3(X18,X19,X20),X20)|(esk3_3(X18,X19,X20)=X18|esk3_3(X18,X19,X20)=X19)|X20=k2_tarski(X18,X19)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_25, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_26, plain, (r1_xboole_0(k1_relat_1(X1),X2)|k9_relat_1(X1,X2)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (k11_relat_1(X1,X2)=k9_relat_1(X1,k1_tarski(X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, negated_conjecture, (r1_xboole_0(X1,k11_relat_1(esk2_0,X2))|r2_hidden(X2,k1_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
fof(c_0_30, plain, ![X35]:(X35=k1_xboole_0|r2_hidden(esk7_1(X35),X35)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_31, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, plain, (k9_relat_1(X1,X2)!=k1_xboole_0|~r2_hidden(X3,k1_relat_1(X1))|~r2_hidden(X3,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_33, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(esk2_0))|k11_relat_1(esk2_0,esk1_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_34, plain, (k11_relat_1(X1,X2)=k9_relat_1(X1,k2_tarski(X2,X2))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_35, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_36, negated_conjecture, (r2_hidden(X1,k1_relat_1(esk2_0))|~r2_hidden(X2,k11_relat_1(esk2_0,X1))|~r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_25, c_0_29])).
cnf(c_0_37, plain, (X1=k1_xboole_0|r2_hidden(esk7_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X3,X1)), inference(er,[status(thm)],[c_0_31])).
cnf(c_0_39, negated_conjecture, (k11_relat_1(esk2_0,esk1_0)!=k1_xboole_0|k9_relat_1(esk2_0,X1)!=k1_xboole_0|~r2_hidden(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_23])])).
cnf(c_0_40, negated_conjecture, (k11_relat_1(esk2_0,X1)=k9_relat_1(esk2_0,k2_tarski(X1,X1))), inference(spm,[status(thm)],[c_0_34, c_0_23])).
cnf(c_0_41, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X1,X3)), inference(er,[status(thm)],[c_0_35])).
cnf(c_0_42, negated_conjecture, (k11_relat_1(esk2_0,esk1_0)=k1_xboole_0|~r2_hidden(esk1_0,k1_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_43, negated_conjecture, (k11_relat_1(esk2_0,X1)=k1_xboole_0|r2_hidden(X1,k1_relat_1(esk2_0))|~r2_hidden(esk7_1(k11_relat_1(esk2_0,X1)),X2)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_44, plain, (r2_hidden(X1,k2_tarski(X2,X1))), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_45, negated_conjecture, (k9_relat_1(esk2_0,k2_tarski(esk1_0,esk1_0))!=k1_xboole_0|k9_relat_1(esk2_0,X1)!=k1_xboole_0|~r2_hidden(esk1_0,X1)), inference(rw,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_46, plain, (r2_hidden(X1,k2_tarski(X1,X2))), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_47, negated_conjecture, (k9_relat_1(esk2_0,k2_tarski(esk1_0,esk1_0))=k1_xboole_0|~r2_hidden(esk1_0,k1_relat_1(esk2_0))), inference(rw,[status(thm)],[c_0_42, c_0_40])).
cnf(c_0_48, negated_conjecture, (k11_relat_1(esk2_0,X1)=k1_xboole_0|r2_hidden(X1,k1_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_49, negated_conjecture, (k9_relat_1(esk2_0,k2_tarski(esk1_0,esk1_0))!=k1_xboole_0|k9_relat_1(esk2_0,k2_tarski(esk1_0,X1))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_50, negated_conjecture, (k9_relat_1(esk2_0,k2_tarski(esk1_0,esk1_0))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_40])])).
cnf(c_0_51, negated_conjecture, (k9_relat_1(esk2_0,k2_tarski(esk1_0,X1))!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_50])])).
cnf(c_0_52, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_50, c_0_51]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 53
# Proof object clause steps            : 34
# Proof object formula steps           : 19
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 22
# Processed clauses                    : 81
# ...of these trivial                  : 1
# ...subsumed                          : 2
# ...remaining for further processing  : 78
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 9
# Generated clauses                    : 126
# ...of the previous two non-trivial   : 125
# Contextual simplify-reflections      : 0
# Paramodulations                      : 118
# Factorizations                       : 0
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 43
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 36
# Current number of unprocessed clauses: 86
# ...number of literals in the above   : 297
# Current number of archived formulas  : 0
# Current number of archived clauses   : 34
# Clause-clause subsumption calls (NU) : 410
# Rec. Clause-clause subsumption calls : 280
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 70
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3400

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28644 pages
