# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t94_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2)), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', t94_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', t48_xboole_1)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', t2_boole)).
fof(idempotence_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', idempotence_k2_xboole_0)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', t3_boole)).
fof(t23_relat_1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>(X3=k1_tarski(k4_tarski(X1,X2))=>(k1_relat_1(X3)=k1_tarski(X1)&k2_relat_1(X3)=k1_tarski(X2)))), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', t23_relat_1)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', t41_xboole_1)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', symmetry_r1_xboole_0)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', t83_xboole_1)).
fof(t65_zfmisc_1, axiom, ![X1, X2]:(k4_xboole_0(X1,k1_tarski(X2))=X1<=>~(r2_hidden(X2,X1))), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', t65_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', t69_enumset1)).
fof(t77_enumset1, axiom, ![X1, X2]:k2_enumset1(X1,X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', t77_enumset1)).
fof(t128_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k1_tarski(X3),X4))<=>(X1=X3&r2_hidden(X2,X4))), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', t128_zfmisc_1)).
fof(t36_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k1_tarski(X1),k2_tarski(X2,X3))=k2_tarski(k4_tarski(X1,X2),k4_tarski(X1,X3))&k2_zfmisc_1(k2_tarski(X1,X2),k1_tarski(X3))=k2_tarski(k4_tarski(X1,X3),k4_tarski(X2,X3))), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', t36_zfmisc_1)).
fof(t129_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,k1_tarski(X4)))<=>(r2_hidden(X1,X3)&X2=X4)), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', t129_zfmisc_1)).
fof(t34_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k1_tarski(X3),k1_tarski(X4)))<=>(X1=X3&X2=X4)), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', t34_zfmisc_1)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', d4_relat_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', d5_relat_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', t3_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', commutativity_k3_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', d3_tarski)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', t36_xboole_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', l32_xboole_1)).
fof(t66_xboole_1, axiom, ![X1]:(~((~(r1_xboole_0(X1,X1))&X1=k1_xboole_0))&~((~(X1=k1_xboole_0)&r1_xboole_0(X1,X1)))), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', t66_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', d10_xboole_0)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/0aa829c0ecaf49fb807454a7b4d2fe7e.E_input', l54_zfmisc_1)).
fof(c_0_26, plain, ![X13, X14]:k2_xboole_0(X13,X14)=k5_xboole_0(k5_xboole_0(X13,X14),k3_xboole_0(X13,X14)), inference(variable_rename,[status(thm)],[t94_xboole_1])).
fof(c_0_27, plain, ![X10, X11]:k4_xboole_0(X10,k4_xboole_0(X10,X11))=k3_xboole_0(X10,X11), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_28, plain, ![X40]:k3_xboole_0(X40,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_29, plain, ![X52]:k2_xboole_0(X52,X52)=X52, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k2_xboole_0])])).
cnf(c_0_30, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_31, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_32, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_33, plain, ![X74]:k4_xboole_0(X74,k1_xboole_0)=X74, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_34, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>(X3=k1_tarski(k4_tarski(X1,X2))=>(k1_relat_1(X3)=k1_tarski(X1)&k2_relat_1(X3)=k1_tarski(X2))))), inference(assume_negation,[status(cth)],[t23_relat_1])).
fof(c_0_35, plain, ![X35, X36, X37]:k4_xboole_0(k4_xboole_0(X35,X36),X37)=k4_xboole_0(X35,k2_xboole_0(X36,X37)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
cnf(c_0_36, plain, (k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k4_xboole_0(X1,k4_xboole_0(X1,X2)))), inference(rw,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_38, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k1_xboole_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_32, c_0_31])).
cnf(c_0_39, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_40, plain, ![X21, X22]:(~r1_xboole_0(X21,X22)|r1_xboole_0(X22,X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
fof(c_0_41, plain, ![X38, X39]:((~r1_xboole_0(X38,X39)|k4_xboole_0(X38,X39)=X38)&(k4_xboole_0(X38,X39)!=X38|r1_xboole_0(X38,X39))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
fof(c_0_42, plain, ![X30, X31]:((k4_xboole_0(X30,k1_tarski(X31))!=X30|~r2_hidden(X31,X30))&(r2_hidden(X31,X30)|k4_xboole_0(X30,k1_tarski(X31))=X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t65_zfmisc_1])])])).
fof(c_0_43, plain, ![X12]:k2_tarski(X12,X12)=k1_tarski(X12), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_44, plain, ![X75, X76]:k2_enumset1(X75,X75,X75,X76)=k2_tarski(X75,X76), inference(variable_rename,[status(thm)],[t77_enumset1])).
fof(c_0_45, negated_conjecture, (v1_relat_1(esk3_0)&(esk3_0=k1_tarski(k4_tarski(esk1_0,esk2_0))&(k1_relat_1(esk3_0)!=k1_tarski(esk1_0)|k2_relat_1(esk3_0)!=k1_tarski(esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_34])])])).
cnf(c_0_46, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_47, plain, (k5_xboole_0(k5_xboole_0(X1,X1),k4_xboole_0(X1,k4_xboole_0(X1,X1)))=X1), inference(rw,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_48, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_49, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_50, plain, (r1_xboole_0(X1,X2)|k4_xboole_0(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_51, plain, (r2_hidden(X1,X2)|k4_xboole_0(X2,k1_tarski(X1))=X2), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_52, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_53, plain, (k2_enumset1(X1,X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_54, negated_conjecture, (esk3_0=k1_tarski(k4_tarski(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_55, plain, ![X17, X18, X19, X20]:(((X17=X19|~r2_hidden(k4_tarski(X17,X18),k2_zfmisc_1(k1_tarski(X19),X20)))&(r2_hidden(X18,X20)|~r2_hidden(k4_tarski(X17,X18),k2_zfmisc_1(k1_tarski(X19),X20))))&(X17!=X19|~r2_hidden(X18,X20)|r2_hidden(k4_tarski(X17,X18),k2_zfmisc_1(k1_tarski(X19),X20)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t128_zfmisc_1])])])).
fof(c_0_56, plain, ![X32, X33, X34]:(k2_zfmisc_1(k1_tarski(X32),k2_tarski(X33,X34))=k2_tarski(k4_tarski(X32,X33),k4_tarski(X32,X34))&k2_zfmisc_1(k2_tarski(X32,X33),k1_tarski(X34))=k2_tarski(k4_tarski(X32,X34),k4_tarski(X33,X34))), inference(variable_rename,[status(thm)],[t36_zfmisc_1])).
cnf(c_0_57, plain, (k4_xboole_0(X1,k1_tarski(X2))!=X1|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_58, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k5_xboole_0(k5_xboole_0(X2,X3),k4_xboole_0(X2,k4_xboole_0(X2,X3))))), inference(rw,[status(thm)],[c_0_46, c_0_37])).
cnf(c_0_59, plain, (k5_xboole_0(k5_xboole_0(X1,X1),X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_48]), c_0_39])).
cnf(c_0_60, plain, (k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_61, plain, (r1_xboole_0(X1,X2)|k4_xboole_0(X2,X1)!=X2), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_62, plain, (k4_xboole_0(X2,k2_enumset1(X1,X1,X1,X1))=X2|r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_52]), c_0_53])).
cnf(c_0_63, negated_conjecture, (esk3_0=k2_enumset1(k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_52]), c_0_53])).
fof(c_0_64, plain, ![X66, X67, X68, X69]:(((r2_hidden(X66,X68)|~r2_hidden(k4_tarski(X66,X67),k2_zfmisc_1(X68,k1_tarski(X69))))&(X67=X69|~r2_hidden(k4_tarski(X66,X67),k2_zfmisc_1(X68,k1_tarski(X69)))))&(~r2_hidden(X66,X68)|X67!=X69|r2_hidden(k4_tarski(X66,X67),k2_zfmisc_1(X68,k1_tarski(X69))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t129_zfmisc_1])])])).
fof(c_0_65, plain, ![X70, X71, X72, X73]:(((X70=X72|~r2_hidden(k4_tarski(X70,X71),k2_zfmisc_1(k1_tarski(X72),k1_tarski(X73))))&(X71=X73|~r2_hidden(k4_tarski(X70,X71),k2_zfmisc_1(k1_tarski(X72),k1_tarski(X73)))))&(X70!=X72|X71!=X73|r2_hidden(k4_tarski(X70,X71),k2_zfmisc_1(k1_tarski(X72),k1_tarski(X73))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t34_zfmisc_1])])])).
cnf(c_0_66, plain, (X1=X2|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k1_tarski(X2),X4))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_67, plain, (k2_zfmisc_1(k2_tarski(X1,X2),k1_tarski(X3))=k2_tarski(k4_tarski(X1,X3),k4_tarski(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_68, plain, (k4_xboole_0(X1,k2_enumset1(X2,X2,X2,X2))!=X1|~r2_hidden(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_52]), c_0_53])).
cnf(c_0_69, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_48]), c_0_39]), c_0_59])).
cnf(c_0_70, plain, (k4_xboole_0(X1,X2)=X1|k4_xboole_0(X2,X1)!=X2), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_71, negated_conjecture, (k4_xboole_0(X1,esk3_0)=X1|r2_hidden(k4_tarski(esk1_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_72, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,k1_tarski(X4)))|~r2_hidden(X1,X2)|X3!=X4), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_73, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k1_tarski(X2),k1_tarski(X4)))|X1!=X2|X3!=X4), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_74, plain, (X1=X2|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k2_enumset1(X2,X2,X2,X2),X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_52]), c_0_53])).
cnf(c_0_75, plain, (k2_zfmisc_1(k2_enumset1(X1,X1,X1,X2),k2_enumset1(X3,X3,X3,X3))=k2_enumset1(k4_tarski(X1,X3),k4_tarski(X1,X3),k4_tarski(X1,X3),k4_tarski(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_52]), c_0_53]), c_0_53]), c_0_53])).
cnf(c_0_76, plain, (~r2_hidden(X1,k4_xboole_0(X2,k2_enumset1(X1,X1,X1,X1)))), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_77, negated_conjecture, (k4_xboole_0(esk3_0,X1)=esk3_0|r2_hidden(k4_tarski(esk1_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
fof(c_0_78, plain, ![X41, X42, X43, X45, X46, X47, X48, X50]:(((~r2_hidden(X43,X42)|r2_hidden(k4_tarski(X43,esk5_3(X41,X42,X43)),X41)|X42!=k1_relat_1(X41))&(~r2_hidden(k4_tarski(X45,X46),X41)|r2_hidden(X45,X42)|X42!=k1_relat_1(X41)))&((~r2_hidden(esk6_2(X47,X48),X48)|~r2_hidden(k4_tarski(esk6_2(X47,X48),X50),X47)|X48=k1_relat_1(X47))&(r2_hidden(esk6_2(X47,X48),X48)|r2_hidden(k4_tarski(esk6_2(X47,X48),esk7_2(X47,X48)),X47)|X48=k1_relat_1(X47)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
cnf(c_0_79, plain, (X1=X2|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,k1_tarski(X2)))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_80, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,k2_enumset1(X4,X4,X4,X4)))|X3!=X4|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_52]), c_0_53])).
fof(c_0_81, plain, ![X55, X56, X57, X59, X60, X61, X62, X64]:(((~r2_hidden(X57,X56)|r2_hidden(k4_tarski(esk8_3(X55,X56,X57),X57),X55)|X56!=k2_relat_1(X55))&(~r2_hidden(k4_tarski(X60,X59),X55)|r2_hidden(X59,X56)|X56!=k2_relat_1(X55)))&((~r2_hidden(esk9_2(X61,X62),X62)|~r2_hidden(k4_tarski(X64,esk9_2(X61,X62)),X61)|X62=k2_relat_1(X61))&(r2_hidden(esk9_2(X61,X62),X62)|r2_hidden(k4_tarski(esk10_2(X61,X62),esk9_2(X61,X62)),X61)|X62=k2_relat_1(X61)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_82, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k2_enumset1(X2,X2,X2,X2),k2_enumset1(X4,X4,X4,X4)))|X3!=X4|X1!=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_52]), c_0_52]), c_0_53]), c_0_53])).
fof(c_0_83, plain, ![X23, X24, X26, X27, X28]:(((r2_hidden(esk4_2(X23,X24),X23)|r1_xboole_0(X23,X24))&(r2_hidden(esk4_2(X23,X24),X24)|r1_xboole_0(X23,X24)))&(~r2_hidden(X28,X26)|~r2_hidden(X28,X27)|~r1_xboole_0(X26,X27))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_84, plain, ![X53, X54]:k3_xboole_0(X53,X54)=k3_xboole_0(X54,X53), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_85, plain, ![X79, X80, X81, X82, X83]:((~r1_tarski(X79,X80)|(~r2_hidden(X81,X79)|r2_hidden(X81,X80)))&((r2_hidden(esk11_2(X82,X83),X82)|r1_tarski(X82,X83))&(~r2_hidden(esk11_2(X82,X83),X83)|r1_tarski(X82,X83)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_86, plain, ![X15, X16]:r1_tarski(k4_xboole_0(X15,X16),X15), inference(variable_rename,[status(thm)],[t36_xboole_1])).
cnf(c_0_87, plain, (X1=X2|~r2_hidden(k4_tarski(X1,X3),k2_enumset1(k4_tarski(X2,X4),k4_tarski(X2,X4),k4_tarski(X2,X4),k4_tarski(X2,X4)))), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_88, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),k2_enumset1(X1,X1,X1,X1))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_89, plain, (r2_hidden(k4_tarski(X1,esk5_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_90, plain, (X1=X2|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,k2_enumset1(X2,X2,X2,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_52]), c_0_53])).
cnf(c_0_91, negated_conjecture, (X1=k4_tarski(esk1_0,esk2_0)|~r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(esk3_0,X3))), inference(spm,[status(thm)],[c_0_74, c_0_63])).
cnf(c_0_92, plain, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,k2_enumset1(X2,X2,X2,X2)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_80])).
cnf(c_0_93, plain, (r2_hidden(k4_tarski(esk8_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_94, plain, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k2_enumset1(X1,X1,X1,X1),k2_enumset1(X2,X2,X2,X2)))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_82])])).
cnf(c_0_95, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_96, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_97, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_98, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
fof(c_0_99, plain, ![X77, X78]:((k4_xboole_0(X77,X78)!=k1_xboole_0|r1_tarski(X77,X78))&(~r1_tarski(X77,X78)|k4_xboole_0(X77,X78)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_100, negated_conjecture, (esk1_0=X1|~r2_hidden(k4_tarski(X1,X2),esk3_0)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_101, plain, (r2_hidden(k4_tarski(X1,esk5_3(X2,k1_relat_1(X2),X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_89])).
fof(c_0_102, plain, ![X29]:((r1_xboole_0(X29,X29)|X29!=k1_xboole_0)&(X29=k1_xboole_0|~r1_xboole_0(X29,X29))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t66_xboole_1])])])).
cnf(c_0_103, plain, (X1=X2|~r2_hidden(k4_tarski(X3,X1),k2_enumset1(k4_tarski(X4,X2),k4_tarski(X4,X2),k4_tarski(X4,X2),k4_tarski(X5,X2)))), inference(spm,[status(thm)],[c_0_90, c_0_75])).
cnf(c_0_104, negated_conjecture, (X1=k4_tarski(esk1_0,esk2_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_105, plain, (r2_hidden(k4_tarski(esk8_3(X1,k2_relat_1(X1),X2),X2),X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_93])).
cnf(c_0_106, plain, (r2_hidden(k4_tarski(X1,X2),k2_enumset1(k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2)))), inference(rw,[status(thm)],[c_0_94, c_0_75])).
cnf(c_0_107, plain, (k4_xboole_0(X1,X2)!=X1|~r2_hidden(X3,X2)|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_95, c_0_50])).
cnf(c_0_108, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_96, c_0_31]), c_0_31])).
cnf(c_0_109, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_110, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_111, negated_conjecture, (esk1_0=X1|~r2_hidden(X1,k1_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_112, plain, (r2_hidden(esk11_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_113, plain, (r1_xboole_0(X1,X1)|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_114, negated_conjecture, (X1=esk2_0|~r2_hidden(k4_tarski(X2,X1),esk3_0)), inference(spm,[status(thm)],[c_0_103, c_0_63])).
cnf(c_0_115, negated_conjecture, (k4_tarski(esk8_3(esk3_0,k2_relat_1(esk3_0),X1),X1)=k4_tarski(esk1_0,esk2_0)|~r2_hidden(X1,k2_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_116, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_106, c_0_63])).
cnf(c_0_117, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))!=X2|~r2_hidden(X3,k4_xboole_0(X2,X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_109])).
cnf(c_0_118, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_110, c_0_98])).
cnf(c_0_119, plain, (r1_tarski(X1,X2)|~r2_hidden(esk11_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_120, negated_conjecture, (esk11_2(k1_relat_1(esk3_0),X1)=esk1_0|r1_tarski(k1_relat_1(esk3_0),X1)), inference(spm,[status(thm)],[c_0_111, c_0_112])).
cnf(c_0_121, plain, (r1_xboole_0(k1_xboole_0,k1_xboole_0)), inference(er,[status(thm)],[c_0_113])).
cnf(c_0_122, negated_conjecture, (X1=esk2_0|~r2_hidden(X1,k2_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_116])])).
cnf(c_0_123, plain, (k4_xboole_0(X1,X2)!=X1|~r2_hidden(X3,k4_xboole_0(X1,k4_xboole_0(X1,X2)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_39])).
cnf(c_0_124, negated_conjecture, (r1_tarski(k1_relat_1(esk3_0),X1)|~r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_119, c_0_120])).
cnf(c_0_125, plain, (r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_78])).
fof(c_0_126, plain, ![X8, X9]:(((r1_tarski(X8,X9)|X8!=X9)&(r1_tarski(X9,X8)|X8!=X9))&(~r1_tarski(X8,X9)|~r1_tarski(X9,X8)|X8=X9)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_127, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_95, c_0_121])).
cnf(c_0_128, negated_conjecture, (esk11_2(k2_relat_1(esk3_0),X1)=esk2_0|r1_tarski(k2_relat_1(esk3_0),X1)), inference(spm,[status(thm)],[c_0_122, c_0_112])).
cnf(c_0_129, plain, (k4_xboole_0(X1,X2)!=X1|~r2_hidden(X3,k4_xboole_0(X2,k4_xboole_0(X2,X1)))), inference(spm,[status(thm)],[c_0_123, c_0_108])).
cnf(c_0_130, negated_conjecture, (k4_xboole_0(k1_relat_1(esk3_0),X1)=k1_xboole_0|~r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_110, c_0_124])).
cnf(c_0_131, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)), inference(er,[status(thm)],[c_0_125])).
cnf(c_0_132, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_126])).
cnf(c_0_133, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_98, c_0_48])).
cnf(c_0_134, plain, (~r2_hidden(X1,k1_relat_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_127, c_0_101])).
cnf(c_0_135, negated_conjecture, (r1_tarski(k2_relat_1(esk3_0),X1)|~r2_hidden(esk2_0,X1)), inference(spm,[status(thm)],[c_0_119, c_0_128])).
cnf(c_0_136, plain, (r2_hidden(X2,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_137, negated_conjecture, (k4_xboole_0(X1,k1_relat_1(esk3_0))!=X1|~r2_hidden(X2,k1_relat_1(esk3_0))|~r2_hidden(esk1_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130]), c_0_39])).
cnf(c_0_138, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_131, c_0_116])).
cnf(c_0_139, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_132, c_0_133])).
cnf(c_0_140, plain, (r1_tarski(k1_relat_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_134, c_0_112])).
cnf(c_0_141, negated_conjecture, (k4_xboole_0(k2_relat_1(esk3_0),X1)=k1_xboole_0|~r2_hidden(esk2_0,X1)), inference(spm,[status(thm)],[c_0_110, c_0_135])).
cnf(c_0_142, plain, (r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)), inference(er,[status(thm)],[c_0_136])).
cnf(c_0_143, negated_conjecture, (k1_relat_1(esk3_0)!=k1_xboole_0|~r2_hidden(X1,k1_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_130]), c_0_138])])).
cnf(c_0_144, plain, (r2_hidden(esk6_2(X1,X2),X2)|r2_hidden(k4_tarski(esk6_2(X1,X2),esk7_2(X1,X2)),X1)|X2=k1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_145, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_139, c_0_140])).
cnf(c_0_146, negated_conjecture, (k4_xboole_0(X1,k2_relat_1(esk3_0))!=X1|~r2_hidden(X2,k2_relat_1(esk3_0))|~r2_hidden(esk2_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_141]), c_0_39])).
cnf(c_0_147, negated_conjecture, (r2_hidden(esk2_0,k2_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_142, c_0_116])).
cnf(c_0_148, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_70, c_0_69])).
cnf(c_0_149, negated_conjecture, (k1_relat_1(esk3_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_143, c_0_138])).
fof(c_0_150, plain, ![X85, X86, X87, X88]:(((r2_hidden(X85,X87)|~r2_hidden(k4_tarski(X85,X86),k2_zfmisc_1(X87,X88)))&(r2_hidden(X86,X88)|~r2_hidden(k4_tarski(X85,X86),k2_zfmisc_1(X87,X88))))&(~r2_hidden(X85,X87)|~r2_hidden(X86,X88)|r2_hidden(k4_tarski(X85,X86),k2_zfmisc_1(X87,X88)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
cnf(c_0_151, plain, (X1=k1_xboole_0|r2_hidden(esk6_2(k1_xboole_0,X1),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_144]), c_0_145])).
cnf(c_0_152, negated_conjecture, (k2_relat_1(esk3_0)!=k1_xboole_0|~r2_hidden(X1,k2_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_141]), c_0_147])])).
cnf(c_0_153, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_98, c_0_108])).
cnf(c_0_154, negated_conjecture, (~r2_hidden(esk1_0,k4_xboole_0(X1,k1_relat_1(esk3_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_130]), c_0_149])).
cnf(c_0_155, plain, (k4_xboole_0(k2_enumset1(X1,X1,X1,X1),X2)=k1_xboole_0|r2_hidden(X1,k4_xboole_0(k2_enumset1(X1,X1,X1,X1),X2))), inference(spm,[status(thm)],[c_0_118, c_0_62])).
cnf(c_0_156, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,X2))), inference(split_conjunct,[status(thm)],[c_0_150])).
cnf(c_0_157, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_xboole_0|k4_xboole_0(X2,X1)!=X2), inference(spm,[status(thm)],[c_0_129, c_0_151])).
cnf(c_0_158, negated_conjecture, (k2_relat_1(esk3_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_152, c_0_147])).
cnf(c_0_159, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=X2|~r1_tarski(X2,k4_xboole_0(X1,k4_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_132, c_0_153])).
cnf(c_0_160, negated_conjecture, (k4_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),k1_relat_1(esk3_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_154, c_0_155])).
cnf(c_0_161, plain, (r2_hidden(X1,k2_enumset1(X2,X2,X2,X2))|~r2_hidden(k4_tarski(X3,X1),k2_enumset1(k4_tarski(X4,X2),k4_tarski(X4,X2),k4_tarski(X4,X2),k4_tarski(X5,X2)))), inference(spm,[status(thm)],[c_0_156, c_0_75])).
cnf(c_0_162, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_163, negated_conjecture, (k4_xboole_0(X1,k2_relat_1(esk3_0))!=X1|~r2_hidden(esk2_0,X1)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_141]), c_0_39]), c_0_158])).
cnf(c_0_164, negated_conjecture, (k1_relat_1(esk3_0)!=k1_tarski(esk1_0)|k2_relat_1(esk3_0)!=k1_tarski(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_165, negated_conjecture, (k1_relat_1(esk3_0)=k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)|~r1_tarski(k1_relat_1(esk3_0),k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_160]), c_0_39]), c_0_39])).
cnf(c_0_166, plain, (r2_hidden(X1,k2_enumset1(X1,X1,X1,X1))), inference(spm,[status(thm)],[c_0_161, c_0_106])).
cnf(c_0_167, plain, (X1=X2|k4_xboole_0(X2,X1)!=k1_xboole_0|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_132, c_0_162])).
cnf(c_0_168, negated_conjecture, (~r2_hidden(esk2_0,k4_xboole_0(X1,k2_relat_1(esk3_0)))), inference(spm,[status(thm)],[c_0_163, c_0_69])).
cnf(c_0_169, negated_conjecture, (k1_relat_1(esk3_0)!=k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)|k2_relat_1(esk3_0)!=k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_164, c_0_52]), c_0_52]), c_0_53]), c_0_53])).
cnf(c_0_170, negated_conjecture, (k1_relat_1(esk3_0)=k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_124]), c_0_166])])).
cnf(c_0_171, negated_conjecture, (k2_relat_1(esk3_0)=X1|k4_xboole_0(X1,k2_relat_1(esk3_0))!=k1_xboole_0|~r2_hidden(esk2_0,X1)), inference(spm,[status(thm)],[c_0_167, c_0_135])).
cnf(c_0_172, negated_conjecture, (k4_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0),k2_relat_1(esk3_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_168, c_0_155])).
cnf(c_0_173, negated_conjecture, (k2_relat_1(esk3_0)!=k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_169, c_0_170])])).
cnf(c_0_174, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_172]), c_0_166])]), c_0_173]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 175
# Proof object clause steps            : 122
# Proof object formula steps           : 53
# Proof object conjectures             : 42
# Proof object clause conjectures      : 39
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 36
# Proof object initial formulas used   : 26
# Proof object generating inferences   : 61
# Proof object simplifying inferences  : 67
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 26
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 53
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 49
# Processed clauses                    : 9489
# ...of these trivial                  : 113
# ...subsumed                          : 8214
# ...remaining for further processing  : 1162
# Other redundant clauses eliminated   : 1769
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 110
# Backward-rewritten                   : 130
# Generated clauses                    : 245304
# ...of the previous two non-trivial   : 213609
# Contextual simplify-reflections      : 11
# Paramodulations                      : 243520
# Factorizations                       : 12
# Equation resolutions                 : 1769
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 864
#    Positive orientable unit clauses  : 151
#    Positive unorientable unit clauses: 4
#    Negative unit clauses             : 24
#    Non-unit-clauses                  : 685
# Current number of unprocessed clauses: 203332
# ...number of literals in the above   : 478795
# Current number of archived formulas  : 0
# Current number of archived clauses   : 292
# Clause-clause subsumption calls (NU) : 161546
# Rec. Clause-clause subsumption calls : 111372
# Non-unit clause-clause subsumptions  : 2712
# Unit Clause-clause subsumption calls : 3421
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2369
# BW rewrite match successes           : 41
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 9738528

# -------------------------------------------------
# User time                : 4.080 s
# System time              : 0.090 s
# Total time               : 4.170 s
# Maximum resident set size: 205380 pages
