# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d7_binop_1, axiom, ![X1, X2]:(m1_subset_1(X2,X1)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))=>(r3_binop_1(X1,X2,X3)<=>(r1_binop_1(X1,X2,X3)&r2_binop_1(X1,X2,X3))))), file('data/MPTP/proofs/0ad29139a2b243588da9d60005b7a0e8.E_input', d7_binop_1)).
fof(dt_k6_filter_1, axiom, ![X1, X2, X3, X4]:(((v1_funct_1(X3)&(v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))&(v1_funct_1(X4)&(v1_funct_2(X4,k2_zfmisc_1(X2,X2),X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2))))))=>(v1_funct_1(k6_filter_1(X1,X2,X3,X4))&(v1_funct_2(k6_filter_1(X1,X2,X3,X4),k2_zfmisc_1(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X2)),k2_zfmisc_1(X1,X2))&m1_subset_1(k6_filter_1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X2)),k2_zfmisc_1(X1,X2))))))), file('data/MPTP/proofs/0ad29139a2b243588da9d60005b7a0e8.E_input', dt_k6_filter_1)).
fof(t27_filter_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,X1)=>![X4]:(m1_subset_1(X4,X2)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,k2_zfmisc_1(X2,X2),X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))))=>((r3_binop_1(X1,X3,X5)&r3_binop_1(X2,X4,X6))<=>r3_binop_1(k2_zfmisc_1(X1,X2),k1_domain_1(X1,X2,X3,X4),k6_filter_1(X1,X2,X5,X6))))))))), file('data/MPTP/proofs/0ad29139a2b243588da9d60005b7a0e8.E_input', t27_filter_1)).
fof(t26_filter_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,X1)=>![X4]:(m1_subset_1(X4,X2)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,k2_zfmisc_1(X2,X2),X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))))=>((r2_binop_1(X1,X3,X5)&r2_binop_1(X2,X4,X6))<=>r2_binop_1(k2_zfmisc_1(X1,X2),k1_domain_1(X1,X2,X3,X4),k6_filter_1(X1,X2,X5,X6))))))))), file('data/MPTP/proofs/0ad29139a2b243588da9d60005b7a0e8.E_input', t26_filter_1)).
fof(t25_filter_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,X1)=>![X4]:(m1_subset_1(X4,X2)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,k2_zfmisc_1(X2,X2),X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))))=>((r1_binop_1(X1,X3,X5)&r1_binop_1(X2,X4,X6))<=>r1_binop_1(k2_zfmisc_1(X1,X2),k1_domain_1(X1,X2,X3,X4),k6_filter_1(X1,X2,X5,X6))))))))), file('data/MPTP/proofs/0ad29139a2b243588da9d60005b7a0e8.E_input', t25_filter_1)).
fof(dt_k1_domain_1, axiom, ![X1, X2, X3, X4]:((~(v1_xboole_0(X1))&(~(v1_xboole_0(X2))&(m1_subset_1(X3,X1)&m1_subset_1(X4,X2))))=>m1_subset_1(k1_domain_1(X1,X2,X3,X4),k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/0ad29139a2b243588da9d60005b7a0e8.E_input', dt_k1_domain_1)).
fof(c_0_6, plain, ![X7, X8, X9]:(((r1_binop_1(X7,X8,X9)|~r3_binop_1(X7,X8,X9)|(~v1_funct_1(X9)|(~v1_funct_2(X9,k2_zfmisc_1(X7,X7),X7)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X7,X7),X7)))))|~m1_subset_1(X8,X7))&(r2_binop_1(X7,X8,X9)|~r3_binop_1(X7,X8,X9)|(~v1_funct_1(X9)|(~v1_funct_2(X9,k2_zfmisc_1(X7,X7),X7)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X7,X7),X7)))))|~m1_subset_1(X8,X7)))&(~r1_binop_1(X7,X8,X9)|~r2_binop_1(X7,X8,X9)|r3_binop_1(X7,X8,X9)|(~v1_funct_1(X9)|(~v1_funct_2(X9,k2_zfmisc_1(X7,X7),X7)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X7,X7),X7)))))|~m1_subset_1(X8,X7))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_binop_1])])])])).
fof(c_0_7, plain, ![X10, X11, X12, X13]:((v1_funct_1(k6_filter_1(X10,X11,X12,X13))|(~v1_funct_1(X12)|(~v1_funct_2(X12,k2_zfmisc_1(X10,X10),X10)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X10,X10),X10))))|(~v1_funct_1(X13)|(~v1_funct_2(X13,k2_zfmisc_1(X11,X11),X11)|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X11,X11),X11)))))))&((v1_funct_2(k6_filter_1(X10,X11,X12,X13),k2_zfmisc_1(k2_zfmisc_1(X10,X11),k2_zfmisc_1(X10,X11)),k2_zfmisc_1(X10,X11))|(~v1_funct_1(X12)|(~v1_funct_2(X12,k2_zfmisc_1(X10,X10),X10)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X10,X10),X10))))|(~v1_funct_1(X13)|(~v1_funct_2(X13,k2_zfmisc_1(X11,X11),X11)|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X11,X11),X11)))))))&(m1_subset_1(k6_filter_1(X10,X11,X12,X13),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X10,X11),k2_zfmisc_1(X10,X11)),k2_zfmisc_1(X10,X11))))|(~v1_funct_1(X12)|(~v1_funct_2(X12,k2_zfmisc_1(X10,X10),X10)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X10,X10),X10))))|(~v1_funct_1(X13)|(~v1_funct_2(X13,k2_zfmisc_1(X11,X11),X11)|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X11,X11),X11))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k6_filter_1])])])).
fof(c_0_8, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,X1)=>![X4]:(m1_subset_1(X4,X2)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,k2_zfmisc_1(X2,X2),X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))))=>((r3_binop_1(X1,X3,X5)&r3_binop_1(X2,X4,X6))<=>r3_binop_1(k2_zfmisc_1(X1,X2),k1_domain_1(X1,X2,X3,X4),k6_filter_1(X1,X2,X5,X6)))))))))), inference(assume_negation,[status(cth)],[t27_filter_1])).
cnf(c_0_9, plain, (r2_binop_1(X1,X2,X3)|~r3_binop_1(X1,X2,X3)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (m1_subset_1(k6_filter_1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X2)),k2_zfmisc_1(X1,X2))))|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~v1_funct_1(X4)|~v1_funct_2(X4,k2_zfmisc_1(X2,X2),X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (v1_funct_1(k6_filter_1(X1,X2,X3,X4))|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~v1_funct_1(X4)|~v1_funct_2(X4,k2_zfmisc_1(X2,X2),X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (v1_funct_2(k6_filter_1(X1,X2,X3,X4),k2_zfmisc_1(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X2)),k2_zfmisc_1(X1,X2))|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~v1_funct_1(X4)|~v1_funct_2(X4,k2_zfmisc_1(X2,X2),X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_13, negated_conjecture, (~v1_xboole_0(esk1_0)&(~v1_xboole_0(esk2_0)&(m1_subset_1(esk3_0,esk1_0)&(m1_subset_1(esk4_0,esk2_0)&((v1_funct_1(esk5_0)&(v1_funct_2(esk5_0,k2_zfmisc_1(esk1_0,esk1_0),esk1_0)&m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0),esk1_0)))))&((v1_funct_1(esk6_0)&(v1_funct_2(esk6_0,k2_zfmisc_1(esk2_0,esk2_0),esk2_0)&m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk2_0),esk2_0)))))&((~r3_binop_1(esk1_0,esk3_0,esk5_0)|~r3_binop_1(esk2_0,esk4_0,esk6_0)|~r3_binop_1(k2_zfmisc_1(esk1_0,esk2_0),k1_domain_1(esk1_0,esk2_0,esk3_0,esk4_0),k6_filter_1(esk1_0,esk2_0,esk5_0,esk6_0)))&((r3_binop_1(esk1_0,esk3_0,esk5_0)|r3_binop_1(k2_zfmisc_1(esk1_0,esk2_0),k1_domain_1(esk1_0,esk2_0,esk3_0,esk4_0),k6_filter_1(esk1_0,esk2_0,esk5_0,esk6_0)))&(r3_binop_1(esk2_0,esk4_0,esk6_0)|r3_binop_1(k2_zfmisc_1(esk1_0,esk2_0),k1_domain_1(esk1_0,esk2_0,esk3_0,esk4_0),k6_filter_1(esk1_0,esk2_0,esk5_0,esk6_0))))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])).
cnf(c_0_14, plain, (r1_binop_1(X1,X2,X3)|~r3_binop_1(X1,X2,X3)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15, plain, (r2_binop_1(k2_zfmisc_1(X1,X2),X3,k6_filter_1(X1,X2,X4,X5))|~r3_binop_1(k2_zfmisc_1(X1,X2),X3,k6_filter_1(X1,X2,X4,X5))|~v1_funct_2(X5,k2_zfmisc_1(X2,X2),X2)|~v1_funct_2(X4,k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X3,k2_zfmisc_1(X1,X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11]), c_0_12])).
cnf(c_0_16, negated_conjecture, (m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk2_0),esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, negated_conjecture, (v1_funct_2(esk6_0,k2_zfmisc_1(esk2_0,esk2_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, negated_conjecture, (v1_funct_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (r1_binop_1(k2_zfmisc_1(X1,X2),X3,k6_filter_1(X1,X2,X4,X5))|~r3_binop_1(k2_zfmisc_1(X1,X2),X3,k6_filter_1(X1,X2,X4,X5))|~v1_funct_2(X5,k2_zfmisc_1(X2,X2),X2)|~v1_funct_2(X4,k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X3,k2_zfmisc_1(X1,X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_10]), c_0_11]), c_0_12])).
fof(c_0_20, plain, ![X20, X21, X22, X23, X24, X25]:((~r2_binop_1(X20,X22,X24)|~r2_binop_1(X21,X23,X25)|r2_binop_1(k2_zfmisc_1(X20,X21),k1_domain_1(X20,X21,X22,X23),k6_filter_1(X20,X21,X24,X25))|(~v1_funct_1(X25)|(~v1_funct_2(X25,k2_zfmisc_1(X21,X21),X21)|~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X21,X21),X21)))))|(~v1_funct_1(X24)|(~v1_funct_2(X24,k2_zfmisc_1(X20,X20),X20)|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X20,X20),X20)))))|~m1_subset_1(X23,X21)|~m1_subset_1(X22,X20)|v1_xboole_0(X21)|v1_xboole_0(X20))&((r2_binop_1(X20,X22,X24)|~r2_binop_1(k2_zfmisc_1(X20,X21),k1_domain_1(X20,X21,X22,X23),k6_filter_1(X20,X21,X24,X25))|(~v1_funct_1(X25)|(~v1_funct_2(X25,k2_zfmisc_1(X21,X21),X21)|~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X21,X21),X21)))))|(~v1_funct_1(X24)|(~v1_funct_2(X24,k2_zfmisc_1(X20,X20),X20)|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X20,X20),X20)))))|~m1_subset_1(X23,X21)|~m1_subset_1(X22,X20)|v1_xboole_0(X21)|v1_xboole_0(X20))&(r2_binop_1(X21,X23,X25)|~r2_binop_1(k2_zfmisc_1(X20,X21),k1_domain_1(X20,X21,X22,X23),k6_filter_1(X20,X21,X24,X25))|(~v1_funct_1(X25)|(~v1_funct_2(X25,k2_zfmisc_1(X21,X21),X21)|~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X21,X21),X21)))))|(~v1_funct_1(X24)|(~v1_funct_2(X24,k2_zfmisc_1(X20,X20),X20)|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X20,X20),X20)))))|~m1_subset_1(X23,X21)|~m1_subset_1(X22,X20)|v1_xboole_0(X21)|v1_xboole_0(X20)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t26_filter_1])])])])])).
cnf(c_0_21, negated_conjecture, (r2_binop_1(k2_zfmisc_1(X1,esk2_0),X2,k6_filter_1(X1,esk2_0,X3,esk6_0))|~r3_binop_1(k2_zfmisc_1(X1,esk2_0),X2,k6_filter_1(X1,esk2_0,X3,esk6_0))|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X2,k2_zfmisc_1(X1,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_18])])).
cnf(c_0_22, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0),esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23, negated_conjecture, (v1_funct_2(esk5_0,k2_zfmisc_1(esk1_0,esk1_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_24, negated_conjecture, (v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_25, plain, ![X26, X27, X28, X29, X30, X31]:((~r1_binop_1(X26,X28,X30)|~r1_binop_1(X27,X29,X31)|r1_binop_1(k2_zfmisc_1(X26,X27),k1_domain_1(X26,X27,X28,X29),k6_filter_1(X26,X27,X30,X31))|(~v1_funct_1(X31)|(~v1_funct_2(X31,k2_zfmisc_1(X27,X27),X27)|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X27,X27),X27)))))|(~v1_funct_1(X30)|(~v1_funct_2(X30,k2_zfmisc_1(X26,X26),X26)|~m1_subset_1(X30,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X26,X26),X26)))))|~m1_subset_1(X29,X27)|~m1_subset_1(X28,X26)|v1_xboole_0(X27)|v1_xboole_0(X26))&((r1_binop_1(X26,X28,X30)|~r1_binop_1(k2_zfmisc_1(X26,X27),k1_domain_1(X26,X27,X28,X29),k6_filter_1(X26,X27,X30,X31))|(~v1_funct_1(X31)|(~v1_funct_2(X31,k2_zfmisc_1(X27,X27),X27)|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X27,X27),X27)))))|(~v1_funct_1(X30)|(~v1_funct_2(X30,k2_zfmisc_1(X26,X26),X26)|~m1_subset_1(X30,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X26,X26),X26)))))|~m1_subset_1(X29,X27)|~m1_subset_1(X28,X26)|v1_xboole_0(X27)|v1_xboole_0(X26))&(r1_binop_1(X27,X29,X31)|~r1_binop_1(k2_zfmisc_1(X26,X27),k1_domain_1(X26,X27,X28,X29),k6_filter_1(X26,X27,X30,X31))|(~v1_funct_1(X31)|(~v1_funct_2(X31,k2_zfmisc_1(X27,X27),X27)|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X27,X27),X27)))))|(~v1_funct_1(X30)|(~v1_funct_2(X30,k2_zfmisc_1(X26,X26),X26)|~m1_subset_1(X30,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X26,X26),X26)))))|~m1_subset_1(X29,X27)|~m1_subset_1(X28,X26)|v1_xboole_0(X27)|v1_xboole_0(X26)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t25_filter_1])])])])])).
cnf(c_0_26, negated_conjecture, (r1_binop_1(k2_zfmisc_1(X1,esk2_0),X2,k6_filter_1(X1,esk2_0,X3,esk6_0))|~r3_binop_1(k2_zfmisc_1(X1,esk2_0),X2,k6_filter_1(X1,esk2_0,X3,esk6_0))|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X2,k2_zfmisc_1(X1,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_16]), c_0_17]), c_0_18])])).
cnf(c_0_27, plain, (r2_binop_1(X1,X2,X3)|v1_xboole_0(X1)|v1_xboole_0(X4)|~r2_binop_1(k2_zfmisc_1(X4,X1),k1_domain_1(X4,X1,X5,X2),k6_filter_1(X4,X1,X6,X3))|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~v1_funct_1(X6)|~v1_funct_2(X6,k2_zfmisc_1(X4,X4),X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X4),X4)))|~m1_subset_1(X2,X1)|~m1_subset_1(X5,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, negated_conjecture, (r2_binop_1(k2_zfmisc_1(esk1_0,esk2_0),X1,k6_filter_1(esk1_0,esk2_0,esk5_0,esk6_0))|~r3_binop_1(k2_zfmisc_1(esk1_0,esk2_0),X1,k6_filter_1(esk1_0,esk2_0,esk5_0,esk6_0))|~m1_subset_1(X1,k2_zfmisc_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24])])).
cnf(c_0_29, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_30, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_31, plain, (r1_binop_1(X1,X2,X3)|v1_xboole_0(X1)|v1_xboole_0(X4)|~r1_binop_1(k2_zfmisc_1(X4,X1),k1_domain_1(X4,X1,X5,X2),k6_filter_1(X4,X1,X6,X3))|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~v1_funct_1(X6)|~v1_funct_2(X6,k2_zfmisc_1(X4,X4),X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X4),X4)))|~m1_subset_1(X2,X1)|~m1_subset_1(X5,X4)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, negated_conjecture, (r1_binop_1(k2_zfmisc_1(esk1_0,esk2_0),X1,k6_filter_1(esk1_0,esk2_0,esk5_0,esk6_0))|~r3_binop_1(k2_zfmisc_1(esk1_0,esk2_0),X1,k6_filter_1(esk1_0,esk2_0,esk5_0,esk6_0))|~m1_subset_1(X1,k2_zfmisc_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_22]), c_0_23]), c_0_24])])).
cnf(c_0_33, negated_conjecture, (r2_binop_1(esk2_0,X1,esk6_0)|~r3_binop_1(k2_zfmisc_1(esk1_0,esk2_0),k1_domain_1(esk1_0,esk2_0,X2,X1),k6_filter_1(esk1_0,esk2_0,esk5_0,esk6_0))|~m1_subset_1(k1_domain_1(esk1_0,esk2_0,X2,X1),k2_zfmisc_1(esk1_0,esk2_0))|~m1_subset_1(X2,esk1_0)|~m1_subset_1(X1,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_23]), c_0_17]), c_0_24]), c_0_18]), c_0_22]), c_0_16])]), c_0_29]), c_0_30])).
cnf(c_0_34, negated_conjecture, (r3_binop_1(esk2_0,esk4_0,esk6_0)|r3_binop_1(k2_zfmisc_1(esk1_0,esk2_0),k1_domain_1(esk1_0,esk2_0,esk3_0,esk4_0),k6_filter_1(esk1_0,esk2_0,esk5_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_35, negated_conjecture, (m1_subset_1(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_36, negated_conjecture, (m1_subset_1(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_37, plain, ![X32, X33, X34, X35]:(v1_xboole_0(X32)|(v1_xboole_0(X33)|(~m1_subset_1(X34,X32)|~m1_subset_1(X35,X33)))|m1_subset_1(k1_domain_1(X32,X33,X34,X35),k2_zfmisc_1(X32,X33))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_domain_1])])])).
cnf(c_0_38, negated_conjecture, (r1_binop_1(esk2_0,X1,esk6_0)|~r3_binop_1(k2_zfmisc_1(esk1_0,esk2_0),k1_domain_1(esk1_0,esk2_0,X2,X1),k6_filter_1(esk1_0,esk2_0,esk5_0,esk6_0))|~m1_subset_1(k1_domain_1(esk1_0,esk2_0,X2,X1),k2_zfmisc_1(esk1_0,esk2_0))|~m1_subset_1(X2,esk1_0)|~m1_subset_1(X1,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_23]), c_0_17]), c_0_24]), c_0_18]), c_0_22]), c_0_16])]), c_0_29]), c_0_30])).
cnf(c_0_39, plain, (r2_binop_1(X1,X2,X3)|v1_xboole_0(X4)|v1_xboole_0(X1)|~r2_binop_1(k2_zfmisc_1(X1,X4),k1_domain_1(X1,X4,X2,X5),k6_filter_1(X1,X4,X3,X6))|~v1_funct_1(X6)|~v1_funct_2(X6,k2_zfmisc_1(X4,X4),X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X4),X4)))|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X5,X4)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_40, plain, (r1_binop_1(X1,X2,X3)|v1_xboole_0(X4)|v1_xboole_0(X1)|~r1_binop_1(k2_zfmisc_1(X1,X4),k1_domain_1(X1,X4,X2,X5),k6_filter_1(X1,X4,X3,X6))|~v1_funct_1(X6)|~v1_funct_2(X6,k2_zfmisc_1(X4,X4),X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X4),X4)))|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X5,X4)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_41, plain, (r3_binop_1(X1,X2,X3)|~r1_binop_1(X1,X2,X3)|~r2_binop_1(X1,X2,X3)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_42, negated_conjecture, (r2_binop_1(esk2_0,esk4_0,esk6_0)|r3_binop_1(esk2_0,esk4_0,esk6_0)|~m1_subset_1(k1_domain_1(esk1_0,esk2_0,esk3_0,esk4_0),k2_zfmisc_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35]), c_0_36])])).
cnf(c_0_43, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(k1_domain_1(X1,X2,X3,X4),k2_zfmisc_1(X1,X2))|~m1_subset_1(X3,X1)|~m1_subset_1(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_44, negated_conjecture, (r1_binop_1(esk2_0,esk4_0,esk6_0)|r3_binop_1(esk2_0,esk4_0,esk6_0)|~m1_subset_1(k1_domain_1(esk1_0,esk2_0,esk3_0,esk4_0),k2_zfmisc_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_34]), c_0_35]), c_0_36])])).
cnf(c_0_45, negated_conjecture, (r2_binop_1(esk1_0,X1,esk5_0)|~r3_binop_1(k2_zfmisc_1(esk1_0,esk2_0),k1_domain_1(esk1_0,esk2_0,X1,X2),k6_filter_1(esk1_0,esk2_0,esk5_0,esk6_0))|~m1_subset_1(k1_domain_1(esk1_0,esk2_0,X1,X2),k2_zfmisc_1(esk1_0,esk2_0))|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_28]), c_0_17]), c_0_23]), c_0_18]), c_0_24]), c_0_16]), c_0_22])]), c_0_30]), c_0_29])).
cnf(c_0_46, negated_conjecture, (r3_binop_1(esk1_0,esk3_0,esk5_0)|r3_binop_1(k2_zfmisc_1(esk1_0,esk2_0),k1_domain_1(esk1_0,esk2_0,esk3_0,esk4_0),k6_filter_1(esk1_0,esk2_0,esk5_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_47, negated_conjecture, (r1_binop_1(esk1_0,X1,esk5_0)|~r3_binop_1(k2_zfmisc_1(esk1_0,esk2_0),k1_domain_1(esk1_0,esk2_0,X1,X2),k6_filter_1(esk1_0,esk2_0,esk5_0,esk6_0))|~m1_subset_1(k1_domain_1(esk1_0,esk2_0,X1,X2),k2_zfmisc_1(esk1_0,esk2_0))|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_32]), c_0_17]), c_0_23]), c_0_18]), c_0_24]), c_0_16]), c_0_22])]), c_0_30]), c_0_29])).
cnf(c_0_48, negated_conjecture, (r3_binop_1(esk2_0,X1,esk6_0)|~r2_binop_1(esk2_0,X1,esk6_0)|~r1_binop_1(esk2_0,X1,esk6_0)|~m1_subset_1(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_16]), c_0_17]), c_0_18])])).
cnf(c_0_49, negated_conjecture, (r2_binop_1(esk2_0,esk4_0,esk6_0)|r3_binop_1(esk2_0,esk4_0,esk6_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_36]), c_0_35])]), c_0_30]), c_0_29])).
cnf(c_0_50, negated_conjecture, (r1_binop_1(esk2_0,esk4_0,esk6_0)|r3_binop_1(esk2_0,esk4_0,esk6_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_43]), c_0_36]), c_0_35])]), c_0_30]), c_0_29])).
cnf(c_0_51, negated_conjecture, (r2_binop_1(esk1_0,esk3_0,esk5_0)|r3_binop_1(esk1_0,esk3_0,esk5_0)|~m1_subset_1(k1_domain_1(esk1_0,esk2_0,esk3_0,esk4_0),k2_zfmisc_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_36]), c_0_35])])).
cnf(c_0_52, negated_conjecture, (r1_binop_1(esk1_0,esk3_0,esk5_0)|r3_binop_1(esk1_0,esk3_0,esk5_0)|~m1_subset_1(k1_domain_1(esk1_0,esk2_0,esk3_0,esk4_0),k2_zfmisc_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_46]), c_0_36]), c_0_35])])).
cnf(c_0_53, negated_conjecture, (~r3_binop_1(esk1_0,esk3_0,esk5_0)|~r3_binop_1(esk2_0,esk4_0,esk6_0)|~r3_binop_1(k2_zfmisc_1(esk1_0,esk2_0),k1_domain_1(esk1_0,esk2_0,esk3_0,esk4_0),k6_filter_1(esk1_0,esk2_0,esk5_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_54, negated_conjecture, (r3_binop_1(esk2_0,esk4_0,esk6_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_36])]), c_0_50])).
cnf(c_0_55, negated_conjecture, (r3_binop_1(esk1_0,X1,esk5_0)|~r2_binop_1(esk1_0,X1,esk5_0)|~r1_binop_1(esk1_0,X1,esk5_0)|~m1_subset_1(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_22]), c_0_23]), c_0_24])])).
cnf(c_0_56, negated_conjecture, (r2_binop_1(esk1_0,esk3_0,esk5_0)|r3_binop_1(esk1_0,esk3_0,esk5_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_43]), c_0_36]), c_0_35])]), c_0_30]), c_0_29])).
cnf(c_0_57, negated_conjecture, (r1_binop_1(esk1_0,esk3_0,esk5_0)|r3_binop_1(esk1_0,esk3_0,esk5_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_43]), c_0_36]), c_0_35])]), c_0_30]), c_0_29])).
cnf(c_0_58, plain, (r3_binop_1(k2_zfmisc_1(X1,X2),X3,k6_filter_1(X1,X2,X4,X5))|~r2_binop_1(k2_zfmisc_1(X1,X2),X3,k6_filter_1(X1,X2,X4,X5))|~r1_binop_1(k2_zfmisc_1(X1,X2),X3,k6_filter_1(X1,X2,X4,X5))|~v1_funct_2(X5,k2_zfmisc_1(X2,X2),X2)|~v1_funct_2(X4,k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X3,k2_zfmisc_1(X1,X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_10]), c_0_11]), c_0_12])).
cnf(c_0_59, plain, (r2_binop_1(k2_zfmisc_1(X1,X4),k1_domain_1(X1,X4,X2,X5),k6_filter_1(X1,X4,X3,X6))|v1_xboole_0(X4)|v1_xboole_0(X1)|~r2_binop_1(X1,X2,X3)|~r2_binop_1(X4,X5,X6)|~v1_funct_1(X6)|~v1_funct_2(X6,k2_zfmisc_1(X4,X4),X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X4),X4)))|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X5,X4)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_60, negated_conjecture, (~r3_binop_1(k2_zfmisc_1(esk1_0,esk2_0),k1_domain_1(esk1_0,esk2_0,esk3_0,esk4_0),k6_filter_1(esk1_0,esk2_0,esk5_0,esk6_0))|~r3_binop_1(esk1_0,esk3_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_54])])).
cnf(c_0_61, negated_conjecture, (r3_binop_1(esk1_0,esk3_0,esk5_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_35])]), c_0_57])).
cnf(c_0_62, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|r3_binop_1(k2_zfmisc_1(X1,X2),k1_domain_1(X1,X2,X3,X4),k6_filter_1(X1,X2,X5,X6))|~r2_binop_1(X2,X4,X6)|~r2_binop_1(X1,X3,X5)|~r1_binop_1(k2_zfmisc_1(X1,X2),k1_domain_1(X1,X2,X3,X4),k6_filter_1(X1,X2,X5,X6))|~v1_funct_2(X6,k2_zfmisc_1(X2,X2),X2)|~v1_funct_2(X5,k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_43])).
cnf(c_0_63, plain, (r1_binop_1(k2_zfmisc_1(X1,X4),k1_domain_1(X1,X4,X2,X5),k6_filter_1(X1,X4,X3,X6))|v1_xboole_0(X4)|v1_xboole_0(X1)|~r1_binop_1(X1,X2,X3)|~r1_binop_1(X4,X5,X6)|~v1_funct_1(X6)|~v1_funct_2(X6,k2_zfmisc_1(X4,X4),X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X4),X4)))|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X5,X4)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_64, negated_conjecture, (~r3_binop_1(k2_zfmisc_1(esk1_0,esk2_0),k1_domain_1(esk1_0,esk2_0,esk3_0,esk4_0),k6_filter_1(esk1_0,esk2_0,esk5_0,esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_61])])).
cnf(c_0_65, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|r3_binop_1(k2_zfmisc_1(X2,X1),k1_domain_1(X2,X1,X3,X4),k6_filter_1(X2,X1,X5,X6))|~r2_binop_1(X1,X4,X6)|~r2_binop_1(X2,X3,X5)|~r1_binop_1(X1,X4,X6)|~r1_binop_1(X2,X3,X5)|~v1_funct_2(X6,k2_zfmisc_1(X1,X1),X1)|~v1_funct_2(X5,k2_zfmisc_1(X2,X2),X2)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X2)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_66, negated_conjecture, (~r2_binop_1(esk2_0,esk4_0,esk6_0)|~r2_binop_1(esk1_0,esk3_0,esk5_0)|~r1_binop_1(esk2_0,esk4_0,esk6_0)|~r1_binop_1(esk1_0,esk3_0,esk5_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_17]), c_0_23]), c_0_18]), c_0_24]), c_0_16]), c_0_22]), c_0_36]), c_0_35])]), c_0_30]), c_0_29])).
cnf(c_0_67, negated_conjecture, (r2_binop_1(esk2_0,X1,esk6_0)|~r3_binop_1(esk2_0,X1,esk6_0)|~m1_subset_1(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_16]), c_0_17]), c_0_18])])).
cnf(c_0_68, negated_conjecture, (~r2_binop_1(esk1_0,esk3_0,esk5_0)|~r1_binop_1(esk2_0,esk4_0,esk6_0)|~r1_binop_1(esk1_0,esk3_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_54]), c_0_36])])).
cnf(c_0_69, negated_conjecture, (r2_binop_1(esk1_0,X1,esk5_0)|~r3_binop_1(esk1_0,X1,esk5_0)|~m1_subset_1(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_22]), c_0_23]), c_0_24])])).
cnf(c_0_70, negated_conjecture, (~r1_binop_1(esk2_0,esk4_0,esk6_0)|~r1_binop_1(esk1_0,esk3_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_61]), c_0_35])])).
cnf(c_0_71, negated_conjecture, (r1_binop_1(esk2_0,X1,esk6_0)|~r3_binop_1(esk2_0,X1,esk6_0)|~m1_subset_1(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_16]), c_0_17]), c_0_18])])).
cnf(c_0_72, negated_conjecture, (~r1_binop_1(esk1_0,esk3_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_54]), c_0_36])])).
cnf(c_0_73, negated_conjecture, (r1_binop_1(esk1_0,X1,esk5_0)|~r3_binop_1(esk1_0,X1,esk5_0)|~m1_subset_1(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_22]), c_0_23]), c_0_24])])).
cnf(c_0_74, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_61]), c_0_35])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 75
# Proof object clause steps            : 62
# Proof object formula steps           : 13
# Proof object conjectures             : 47
# Proof object clause conjectures      : 44
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 34
# Proof object simplifying inferences  : 138
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 26
# Processed clauses                    : 180
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 180
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 10
# Backward-rewritten                   : 11
# Generated clauses                    : 181
# ...of the previous two non-trivial   : 140
# Contextual simplify-reflections      : 64
# Paramodulations                      : 181
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 133
#    Positive orientable unit clauses  : 30
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 99
# Current number of unprocessed clauses: 12
# ...number of literals in the above   : 160
# Current number of archived formulas  : 0
# Current number of archived clauses   : 47
# Clause-clause subsumption calls (NU) : 3817
# Rec. Clause-clause subsumption calls : 836
# Non-unit clause-clause subsumptions  : 73
# Unit Clause-clause subsumption calls : 87
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 66
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 15646

# -------------------------------------------------
# User time                : 0.018 s
# System time              : 0.003 s
# Total time               : 0.021 s
# Maximum resident set size: 28668 pages
