# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t67_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>((r1_yellow_0(X1,k7_domain_1(u1_struct_0(X2),X3,X4))&r2_hidden(k1_yellow_0(X1,k7_domain_1(u1_struct_0(X2),X3,X4)),u1_struct_0(X2)))=>(r1_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))&k1_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))=k1_yellow_0(X1,k7_domain_1(u1_struct_0(X2),X3,X4)))))))), file('data/MPTP/proofs/0adb371f6fea4925a8ed3ddc24340c1c.E_input', t67_yellow_0)).
fof(t62_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,X1)=>![X3]:(m1_subset_1(X3,X1)=>![X4]:(m1_subset_1(X4,X1)=>![X5]:(m1_subset_1(X5,X1)=>![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X1)=>![X8]:(m1_subset_1(X8,X1)=>![X9]:(m1_subset_1(X9,X1)=>(~(X1=k1_xboole_0)=>m1_subset_1(k6_enumset1(X2,X3,X4,X5,X6,X7,X8,X9),k1_zfmisc_1(X1))))))))))), file('data/MPTP/proofs/0adb371f6fea4925a8ed3ddc24340c1c.E_input', t62_subset_1)).
fof(redefinition_k7_domain_1, axiom, ![X1, X2, X3]:((~(v1_xboole_0(X1))&(m1_subset_1(X2,X1)&m1_subset_1(X3,X1)))=>k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)), file('data/MPTP/proofs/0adb371f6fea4925a8ed3ddc24340c1c.E_input', redefinition_k7_domain_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/0adb371f6fea4925a8ed3ddc24340c1c.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/0adb371f6fea4925a8ed3ddc24340c1c.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/0adb371f6fea4925a8ed3ddc24340c1c.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/0adb371f6fea4925a8ed3ddc24340c1c.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/0adb371f6fea4925a8ed3ddc24340c1c.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/0adb371f6fea4925a8ed3ddc24340c1c.E_input', t75_enumset1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/0adb371f6fea4925a8ed3ddc24340c1c.E_input', d1_xboole_0)).
fof(t65_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>((r1_yellow_0(X1,X3)&r2_hidden(k1_yellow_0(X1,X3),u1_struct_0(X2)))=>(r1_yellow_0(X2,X3)&k1_yellow_0(X2,X3)=k1_yellow_0(X1,X3)))))), file('data/MPTP/proofs/0adb371f6fea4925a8ed3ddc24340c1c.E_input', t65_yellow_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/0adb371f6fea4925a8ed3ddc24340c1c.E_input', fc1_xboole_0)).
fof(c_0_12, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>((r1_yellow_0(X1,k7_domain_1(u1_struct_0(X2),X3,X4))&r2_hidden(k1_yellow_0(X1,k7_domain_1(u1_struct_0(X2),X3,X4)),u1_struct_0(X2)))=>(r1_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))&k1_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))=k1_yellow_0(X1,k7_domain_1(u1_struct_0(X2),X3,X4))))))))), inference(assume_negation,[status(cth)],[t67_yellow_0])).
fof(c_0_13, plain, ![X39, X40, X41, X42, X43, X44, X45, X46, X47]:(~m1_subset_1(X40,X39)|(~m1_subset_1(X41,X39)|(~m1_subset_1(X42,X39)|(~m1_subset_1(X43,X39)|(~m1_subset_1(X44,X39)|(~m1_subset_1(X45,X39)|(~m1_subset_1(X46,X39)|(~m1_subset_1(X47,X39)|(X39=k1_xboole_0|m1_subset_1(k6_enumset1(X40,X41,X42,X43,X44,X45,X46,X47),k1_zfmisc_1(X39))))))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t62_subset_1])])])])).
fof(c_0_14, negated_conjecture, ((~v2_struct_0(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0)))&((~v2_struct_0(esk2_0)&(v4_yellow_0(esk2_0,esk1_0)&m1_yellow_0(esk2_0,esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(esk2_0))&(m1_subset_1(esk4_0,u1_struct_0(esk2_0))&((r1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))&r2_hidden(k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0)),u1_struct_0(esk2_0)))&(~r1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))|k1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))!=k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])).
cnf(c_0_15, plain, (X2=k1_xboole_0|m1_subset_1(k6_enumset1(X1,X3,X4,X5,X6,X7,X8,X9),k1_zfmisc_1(X2))|~m1_subset_1(X1,X2)|~m1_subset_1(X3,X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X7,X2)|~m1_subset_1(X8,X2)|~m1_subset_1(X9,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_16, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_17, negated_conjecture, (k1_xboole_0=u1_struct_0(esk2_0)|m1_subset_1(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,esk4_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X7,u1_struct_0(esk2_0))|~m1_subset_1(X6,u1_struct_0(esk2_0))|~m1_subset_1(X5,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_18, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, negated_conjecture, (k1_xboole_0=u1_struct_0(esk2_0)|m1_subset_1(k6_enumset1(X1,X2,X3,X4,X5,X6,esk3_0,esk4_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X6,u1_struct_0(esk2_0))|~m1_subset_1(X5,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_20, negated_conjecture, (k1_xboole_0=u1_struct_0(esk2_0)|m1_subset_1(k6_enumset1(X1,X2,X3,X4,X5,esk3_0,esk3_0,esk4_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X5,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_19, c_0_18])).
fof(c_0_21, plain, ![X48, X49, X50]:(v1_xboole_0(X48)|(~m1_subset_1(X49,X48)|~m1_subset_1(X50,X48))|k7_domain_1(X48,X49,X50)=k2_tarski(X49,X50)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k7_domain_1])])])).
fof(c_0_22, plain, ![X34, X35]:k1_enumset1(X34,X34,X35)=k2_tarski(X34,X35), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_23, plain, ![X36, X37, X38]:k2_enumset1(X36,X36,X37,X38)=k1_enumset1(X36,X37,X38), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_24, plain, ![X24, X25, X26, X27]:k3_enumset1(X24,X24,X25,X26,X27)=k2_enumset1(X24,X25,X26,X27), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_25, plain, ![X55, X56, X57, X58, X59]:k4_enumset1(X55,X55,X56,X57,X58,X59)=k3_enumset1(X55,X56,X57,X58,X59), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_26, plain, ![X28, X29, X30, X31, X32, X33]:k5_enumset1(X28,X28,X29,X30,X31,X32,X33)=k4_enumset1(X28,X29,X30,X31,X32,X33), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_27, plain, ![X14, X15, X16, X17, X18, X19, X20]:k6_enumset1(X14,X14,X15,X16,X17,X18,X19,X20)=k5_enumset1(X14,X15,X16,X17,X18,X19,X20), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_28, plain, ![X51, X52, X53]:((~v1_xboole_0(X51)|~r2_hidden(X52,X51))&(r2_hidden(esk5_1(X53),X53)|v1_xboole_0(X53))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_29, negated_conjecture, (k1_xboole_0=u1_struct_0(esk2_0)|m1_subset_1(k6_enumset1(X1,X2,X3,X4,esk3_0,esk3_0,esk3_0,esk4_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X4,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_20, c_0_18])).
cnf(c_0_30, plain, (v1_xboole_0(X1)|k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_33, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, negated_conjecture, (r2_hidden(k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0)),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_39, negated_conjecture, (k1_xboole_0=u1_struct_0(esk2_0)|m1_subset_1(k6_enumset1(X1,X2,X3,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_29, c_0_18])).
cnf(c_0_40, plain, (k7_domain_1(X1,X2,X3)=k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)|v1_xboole_0(X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35]), c_0_36])).
cnf(c_0_41, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
fof(c_0_42, plain, ![X21, X22, X23]:((r1_yellow_0(X22,X23)|(~r1_yellow_0(X21,X23)|~r2_hidden(k1_yellow_0(X21,X23),u1_struct_0(X22)))|~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(X22)))|(v2_struct_0(X22)|(~v4_yellow_0(X22,X21)|~m1_yellow_0(X22,X21)))|(v2_struct_0(X21)|(~v4_orders_2(X21)|~l1_orders_2(X21))))&(k1_yellow_0(X22,X23)=k1_yellow_0(X21,X23)|(~r1_yellow_0(X21,X23)|~r2_hidden(k1_yellow_0(X21,X23),u1_struct_0(X22)))|~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(X22)))|(v2_struct_0(X22)|(~v4_yellow_0(X22,X21)|~m1_yellow_0(X22,X21)))|(v2_struct_0(X21)|(~v4_orders_2(X21)|~l1_orders_2(X21))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t65_yellow_0])])])])])).
cnf(c_0_43, negated_conjecture, (k1_xboole_0=u1_struct_0(esk2_0)|m1_subset_1(k6_enumset1(X1,X2,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_39, c_0_18])).
cnf(c_0_44, negated_conjecture, (k6_enumset1(X1,X1,X1,X1,X1,X1,X1,esk4_0)=k7_domain_1(u1_struct_0(esk2_0),X1,esk4_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_16]), c_0_41])).
cnf(c_0_45, plain, (r1_yellow_0(X1,X2)|v2_struct_0(X1)|v2_struct_0(X3)|~r1_yellow_0(X3,X2)|~r2_hidden(k1_yellow_0(X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v4_yellow_0(X1,X3)|~m1_yellow_0(X1,X3)|~v4_orders_2(X3)|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_46, negated_conjecture, (m1_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_47, negated_conjecture, (v4_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_48, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_49, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_50, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_51, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_52, negated_conjecture, (k1_xboole_0=u1_struct_0(esk2_0)|m1_subset_1(k6_enumset1(X1,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_43, c_0_18])).
cnf(c_0_53, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)=k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_44, c_0_18])).
cnf(c_0_54, plain, (k1_yellow_0(X1,X2)=k1_yellow_0(X3,X2)|v2_struct_0(X1)|v2_struct_0(X3)|~r1_yellow_0(X3,X2)|~r2_hidden(k1_yellow_0(X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v4_yellow_0(X1,X3)|~m1_yellow_0(X1,X3)|~v4_orders_2(X3)|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_55, negated_conjecture, (r1_yellow_0(esk2_0,X1)|~r2_hidden(k1_yellow_0(esk1_0,X1),u1_struct_0(esk2_0))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47]), c_0_48]), c_0_49])]), c_0_50]), c_0_51])).
cnf(c_0_56, negated_conjecture, (r1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_57, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_58, negated_conjecture, (k1_xboole_0=u1_struct_0(esk2_0)|m1_subset_1(k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_18]), c_0_53])).
cnf(c_0_59, negated_conjecture, (k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk2_0,X1)|~r2_hidden(k1_yellow_0(esk1_0,X1),u1_struct_0(esk2_0))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_46]), c_0_47]), c_0_48]), c_0_49])]), c_0_50]), c_0_51])).
cnf(c_0_60, negated_conjecture, (r1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))|~m1_subset_1(k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_38]), c_0_56])])).
cnf(c_0_61, negated_conjecture, (m1_subset_1(k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_41])).
cnf(c_0_62, negated_conjecture, (k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))=k1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))|~m1_subset_1(k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_38]), c_0_56])])).
cnf(c_0_63, negated_conjecture, (~r1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))|k1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))!=k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_64, negated_conjecture, (r1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_61])])).
cnf(c_0_65, negated_conjecture, (k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))=k1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_61])])).
cnf(c_0_66, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_64])]), c_0_65])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 67
# Proof object clause steps            : 43
# Proof object formula steps           : 24
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 33
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 6
# Initial clauses in saturation        : 18
# Processed clauses                    : 592
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 590
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 6
# Generated clauses                    : 35457
# ...of the previous two non-trivial   : 35459
# Contextual simplify-reflections      : 0
# Paramodulations                      : 35457
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 566
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 546
# Current number of unprocessed clauses: 34903
# ...number of literals in the above   : 231328
# Current number of archived formulas  : 0
# Current number of archived clauses   : 30
# Clause-clause subsumption calls (NU) : 248336
# Rec. Clause-clause subsumption calls : 11345
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 516
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2302538

# -------------------------------------------------
# User time                : 0.448 s
# System time              : 0.010 s
# Total time               : 0.458 s
# Maximum resident set size: 43236 pages
