# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t23_waybel_9, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(k9_yellow_6(X1,X2)))=>![X4]:(m1_subset_1(X4,u1_struct_0(k9_yellow_6(X1,X2)))=>m1_subset_1(k2_xboole_0(X3,X4),u1_struct_0(k9_yellow_6(X1,X2))))))), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', t23_waybel_9)).
fof(l20_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k2_xboole_0(k1_tarski(X1),X2),X2)=>r2_hidden(X1,X2)), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', l20_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', t69_enumset1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', l51_zfmisc_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', t12_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', d10_xboole_0)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', d2_subset_1)).
fof(t6_connsp_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(m1_connsp_2(X2,X1,X3)=>r2_hidden(X3,X2))))), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', t6_connsp_2)).
fof(dt_m1_connsp_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,u1_struct_0(X1)))=>![X3]:(m1_connsp_2(X3,X1,X2)=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', dt_m1_connsp_2)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', t7_xboole_1)).
fof(t21_waybel_9, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(k9_yellow_6(X1,X2)))=>m1_connsp_2(X3,X1,X2)))), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', t21_waybel_9)).
fof(t11_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k2_xboole_0(X1,X2),X3)=>r1_tarski(X1,X3)), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', t11_xboole_1)).
fof(l22_zfmisc_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>k2_xboole_0(k1_tarski(X1),X2)=X2), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', l22_zfmisc_1)).
fof(t39_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X3,u1_struct_0(k9_yellow_6(X1,X2)))<=>(r2_hidden(X2,X3)&v3_pre_topc(X3,X1)))))), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', t39_yellow_6)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', t4_subset)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', t1_subset)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', d1_xboole_0)).
fof(fc7_tops_1, axiom, ![X1, X2, X3]:(((v2_pre_topc(X1)&l1_pre_topc(X1))&((v3_pre_topc(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(v3_pre_topc(X3,X1)&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))))=>v3_pre_topc(k2_xboole_0(X2,X3),X1)), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', fc7_tops_1)).
fof(redefinition_k4_subset_1, axiom, ![X1, X2, X3]:((m1_subset_1(X2,k1_zfmisc_1(X1))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>k4_subset_1(X1,X2,X3)=k2_xboole_0(X2,X3)), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', redefinition_k4_subset_1)).
fof(dt_k4_subset_1, axiom, ![X1, X2, X3]:((m1_subset_1(X2,k1_zfmisc_1(X1))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>m1_subset_1(k4_subset_1(X1,X2,X3),k1_zfmisc_1(X1))), file('data/MPTP/proofs/0add73e37bde4ef5a21acabf9b66724e.E_input', dt_k4_subset_1)).
fof(c_0_20, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(k9_yellow_6(X1,X2)))=>![X4]:(m1_subset_1(X4,u1_struct_0(k9_yellow_6(X1,X2)))=>m1_subset_1(k2_xboole_0(X3,X4),u1_struct_0(k9_yellow_6(X1,X2)))))))), inference(assume_negation,[status(cth)],[t23_waybel_9])).
fof(c_0_21, plain, ![X35, X36]:(~r1_tarski(k2_xboole_0(k1_tarski(X35),X36),X36)|r2_hidden(X35,X36)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l20_zfmisc_1])])).
fof(c_0_22, plain, ![X16]:k2_tarski(X16,X16)=k1_tarski(X16), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_23, plain, ![X37, X38]:k3_tarski(k2_tarski(X37,X38))=k2_xboole_0(X37,X38), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_24, plain, ![X30, X31]:(~r1_tarski(X30,X31)|k2_xboole_0(X30,X31)=X31), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
fof(c_0_25, plain, ![X54, X55]:(((r1_tarski(X54,X55)|X54!=X55)&(r1_tarski(X55,X54)|X54!=X55))&(~r1_tarski(X54,X55)|~r1_tarski(X55,X54)|X54=X55)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_26, plain, ![X39, X40]:(((~m1_subset_1(X40,X39)|r2_hidden(X40,X39)|v1_xboole_0(X39))&(~r2_hidden(X40,X39)|m1_subset_1(X40,X39)|v1_xboole_0(X39)))&((~m1_subset_1(X40,X39)|v1_xboole_0(X40)|~v1_xboole_0(X39))&(~v1_xboole_0(X40)|m1_subset_1(X40,X39)|~v1_xboole_0(X39)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_27, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))&(m1_subset_1(esk4_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))&~m1_subset_1(k2_xboole_0(esk3_0,esk4_0),u1_struct_0(k9_yellow_6(esk1_0,esk2_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_20])])])])).
fof(c_0_28, plain, ![X19, X20, X21]:(v2_struct_0(X19)|(~v2_pre_topc(X19)|~l1_pre_topc(X19))|(~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|(~m1_subset_1(X21,u1_struct_0(X19))|(~m1_connsp_2(X20,X19,X21)|r2_hidden(X21,X20))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_connsp_2])])])])).
fof(c_0_29, plain, ![X48, X49, X50]:(v2_struct_0(X48)|(~v2_pre_topc(X48)|~l1_pre_topc(X48))|~m1_subset_1(X49,u1_struct_0(X48))|(~m1_connsp_2(X50,X48,X49)|m1_subset_1(X50,k1_zfmisc_1(u1_struct_0(X48))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_connsp_2])])])])).
cnf(c_0_30, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_xboole_0(k1_tarski(X1),X2),X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_33, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_35, plain, ![X22, X23]:r1_tarski(X22,k2_xboole_0(X22,X23)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_36, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_37, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_38, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_39, plain, ![X27, X28, X29]:(v2_struct_0(X27)|(~v2_pre_topc(X27)|~l1_pre_topc(X27))|(~m1_subset_1(X28,u1_struct_0(X27))|(~m1_subset_1(X29,u1_struct_0(k9_yellow_6(X27,X28)))|m1_connsp_2(X29,X27,X28)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t21_waybel_9])])])])).
cnf(c_0_40, plain, (v2_struct_0(X1)|r2_hidden(X3,X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_connsp_2(X2,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_41, plain, (v2_struct_0(X1)|m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_connsp_2(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_42, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_43, plain, (r2_hidden(X1,X2)|~r1_tarski(k3_tarski(k2_tarski(k2_tarski(X1,X1),X2)),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_31]), c_0_32])).
cnf(c_0_44, plain, (k3_tarski(k2_tarski(X1,X2))=X2|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_33, c_0_32])).
cnf(c_0_45, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_34])).
cnf(c_0_46, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_47, plain, ![X56, X57, X58]:(~r1_tarski(k2_xboole_0(X56,X57),X58)|r1_tarski(X56,X58)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_xboole_1])])).
fof(c_0_48, plain, ![X17, X18]:(~r2_hidden(X17,X18)|k2_xboole_0(k1_tarski(X17),X18)=X18), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l22_zfmisc_1])])).
fof(c_0_49, plain, ![X24, X25, X26]:(((r2_hidden(X25,X26)|~r2_hidden(X26,u1_struct_0(k9_yellow_6(X24,X25)))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X24)))|~m1_subset_1(X25,u1_struct_0(X24))|(v2_struct_0(X24)|(~v2_pre_topc(X24)|~l1_pre_topc(X24))))&(v3_pre_topc(X26,X24)|~r2_hidden(X26,u1_struct_0(k9_yellow_6(X24,X25)))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X24)))|~m1_subset_1(X25,u1_struct_0(X24))|(v2_struct_0(X24)|(~v2_pre_topc(X24)|~l1_pre_topc(X24)))))&(~r2_hidden(X25,X26)|~v3_pre_topc(X26,X24)|r2_hidden(X26,u1_struct_0(k9_yellow_6(X24,X25)))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X24)))|~m1_subset_1(X25,u1_struct_0(X24))|(v2_struct_0(X24)|(~v2_pre_topc(X24)|~l1_pre_topc(X24))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t39_yellow_6])])])])])).
fof(c_0_50, plain, ![X43, X44, X45]:(~r2_hidden(X43,X44)|~m1_subset_1(X44,k1_zfmisc_1(X45))|m1_subset_1(X43,X45)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_51, negated_conjecture, (v1_xboole_0(esk4_0)|~v1_xboole_0(u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk4_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))|v1_xboole_0(u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_38, c_0_37])).
cnf(c_0_53, plain, (v2_struct_0(X1)|m1_connsp_2(X3,X1,X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(k9_yellow_6(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_54, plain, (r2_hidden(X1,X2)|v2_struct_0(X3)|~m1_connsp_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X3))|~l1_pre_topc(X3)|~v2_pre_topc(X3)), inference(csr,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_55, negated_conjecture, (r2_hidden(esk3_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))|v1_xboole_0(u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_38, c_0_42])).
fof(c_0_56, plain, ![X46, X47]:(~r2_hidden(X46,X47)|m1_subset_1(X46,X47)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_57, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_tarski(X1,X1),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])])).
cnf(c_0_58, plain, (r1_tarski(X1,k3_tarski(k2_tarski(X1,X2)))), inference(rw,[status(thm)],[c_0_46, c_0_32])).
cnf(c_0_59, plain, (r1_tarski(X1,X3)|~r1_tarski(k2_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_60, plain, (k2_xboole_0(k1_tarski(X1),X2)=X2|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_61, plain, ![X12, X13, X14]:((~v1_xboole_0(X12)|~r2_hidden(X13,X12))&(r2_hidden(esk5_1(X14),X14)|v1_xboole_0(X14))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_62, plain, (r2_hidden(X2,u1_struct_0(k9_yellow_6(X3,X1)))|v2_struct_0(X3)|~r2_hidden(X1,X2)|~v3_pre_topc(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_63, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_64, plain, (v3_pre_topc(X1,X2)|v2_struct_0(X2)|~r2_hidden(X1,u1_struct_0(k9_yellow_6(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,u1_struct_0(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_65, negated_conjecture, (r2_hidden(esk4_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))|v1_xboole_0(esk4_0)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_66, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_67, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_68, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_69, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_70, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(k9_yellow_6(X2,X3)))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_41, c_0_53])).
cnf(c_0_71, plain, (r2_hidden(X1,X2)|v2_struct_0(X3)|~m1_subset_1(X2,u1_struct_0(k9_yellow_6(X3,X1)))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_pre_topc(X3)|~v2_pre_topc(X3)), inference(spm,[status(thm)],[c_0_54, c_0_53])).
cnf(c_0_72, negated_conjecture, (r2_hidden(esk3_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))|v1_xboole_0(esk4_0)), inference(spm,[status(thm)],[c_0_51, c_0_55])).
cnf(c_0_73, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_74, plain, (r2_hidden(X1,k3_tarski(k2_tarski(k2_tarski(X1,X1),X2)))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_75, plain, (r1_tarski(X1,X3)|~r1_tarski(k3_tarski(k2_tarski(X1,X2)),X3)), inference(rw,[status(thm)],[c_0_59, c_0_32])).
cnf(c_0_76, plain, (k3_tarski(k2_tarski(k2_tarski(X1,X1),X2))=X2|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_31]), c_0_32])).
cnf(c_0_77, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_78, negated_conjecture, (~m1_subset_1(k2_xboole_0(esk3_0,esk4_0),u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_79, plain, (r2_hidden(X1,u1_struct_0(k9_yellow_6(X2,X3)))|v2_struct_0(X2)|~v3_pre_topc(X1,X2)|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[c_0_62, c_0_63])).
fof(c_0_80, plain, ![X51, X52, X53]:(~v2_pre_topc(X51)|~l1_pre_topc(X51)|(~v3_pre_topc(X52,X51)|~m1_subset_1(X52,k1_zfmisc_1(u1_struct_0(X51)))|(~v3_pre_topc(X53,X51)|~m1_subset_1(X53,k1_zfmisc_1(u1_struct_0(X51)))))|v3_pre_topc(k2_xboole_0(X52,X53),X51)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc7_tops_1])])).
cnf(c_0_81, negated_conjecture, (v3_pre_topc(esk4_0,esk1_0)|v1_xboole_0(esk4_0)|~m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_66]), c_0_67]), c_0_68])]), c_0_69])).
cnf(c_0_82, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_37]), c_0_66]), c_0_67]), c_0_68])]), c_0_69])).
cnf(c_0_83, negated_conjecture, (r2_hidden(esk2_0,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_37]), c_0_66]), c_0_67]), c_0_68])]), c_0_69])).
cnf(c_0_84, negated_conjecture, (v3_pre_topc(esk3_0,esk1_0)|v1_xboole_0(esk4_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_72]), c_0_66]), c_0_67]), c_0_68])]), c_0_69])).
cnf(c_0_85, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_42]), c_0_66]), c_0_67]), c_0_68])]), c_0_69])).
fof(c_0_86, plain, ![X9, X10, X11]:(~m1_subset_1(X10,k1_zfmisc_1(X9))|~m1_subset_1(X11,k1_zfmisc_1(X9))|k4_subset_1(X9,X10,X11)=k2_xboole_0(X10,X11)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_subset_1])])).
cnf(c_0_87, plain, (m1_subset_1(X1,k3_tarski(k2_tarski(k2_tarski(X1,X1),X2)))), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_88, plain, (r1_tarski(k2_tarski(X1,X1),X2)|~r1_tarski(X3,X2)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_89, negated_conjecture, (r2_hidden(esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_42]), c_0_66]), c_0_67]), c_0_68])]), c_0_69])).
cnf(c_0_90, plain, (~v1_xboole_0(k3_tarski(k2_tarski(k2_tarski(X1,X1),X2)))), inference(spm,[status(thm)],[c_0_77, c_0_74])).
cnf(c_0_91, negated_conjecture, (~m1_subset_1(k3_tarski(k2_tarski(esk3_0,esk4_0)),u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(rw,[status(thm)],[c_0_78, c_0_32])).
cnf(c_0_92, plain, (m1_subset_1(X1,u1_struct_0(k9_yellow_6(X2,X3)))|v2_struct_0(X2)|~v3_pre_topc(X1,X2)|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_73, c_0_79])).
cnf(c_0_93, plain, (v3_pre_topc(k2_xboole_0(X2,X3),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_pre_topc(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_94, negated_conjecture, (v3_pre_topc(esk4_0,esk1_0)|v1_xboole_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_82])])).
cnf(c_0_95, negated_conjecture, (~v1_xboole_0(esk4_0)), inference(spm,[status(thm)],[c_0_77, c_0_83])).
cnf(c_0_96, negated_conjecture, (v3_pre_topc(esk3_0,esk1_0)|v1_xboole_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_84, c_0_85])])).
fof(c_0_97, plain, ![X32, X33, X34]:(~m1_subset_1(X33,k1_zfmisc_1(X32))|~m1_subset_1(X34,k1_zfmisc_1(X32))|m1_subset_1(k4_subset_1(X32,X33,X34),k1_zfmisc_1(X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_subset_1])])).
cnf(c_0_98, plain, (k4_subset_1(X2,X1,X3)=k2_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_99, plain, (m1_subset_1(X1,X2)|~r1_tarski(k2_tarski(X1,X1),X2)), inference(spm,[status(thm)],[c_0_87, c_0_44])).
cnf(c_0_100, negated_conjecture, (r1_tarski(k2_tarski(esk2_0,esk2_0),X1)|~r1_tarski(esk3_0,X1)), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_101, plain, (~r1_tarski(k2_tarski(X1,X1),X2)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_90, c_0_44])).
cnf(c_0_102, negated_conjecture, (~v3_pre_topc(k3_tarski(k2_tarski(esk3_0,esk4_0)),esk1_0)|~r2_hidden(esk2_0,k3_tarski(k2_tarski(esk3_0,esk4_0)))|~m1_subset_1(k3_tarski(k2_tarski(esk3_0,esk4_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_67]), c_0_68])]), c_0_69])).
cnf(c_0_103, plain, (v3_pre_topc(k3_tarski(k2_tarski(X2,X3)),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v3_pre_topc(X3,X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(rw,[status(thm)],[c_0_93, c_0_32])).
cnf(c_0_104, negated_conjecture, (v3_pre_topc(esk4_0,esk1_0)), inference(sr,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_105, negated_conjecture, (v3_pre_topc(esk3_0,esk1_0)), inference(sr,[status(thm)],[c_0_96, c_0_95])).
cnf(c_0_106, plain, (m1_subset_1(k4_subset_1(X2,X1,X3),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_107, plain, (k4_subset_1(X2,X1,X3)=k3_tarski(k2_tarski(X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_98, c_0_32])).
cnf(c_0_108, negated_conjecture, (m1_subset_1(esk2_0,X1)|~r1_tarski(esk3_0,X1)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_109, negated_conjecture, (~r1_tarski(esk3_0,X1)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_101, c_0_100])).
cnf(c_0_110, negated_conjecture, (~r2_hidden(esk2_0,k3_tarski(k2_tarski(esk3_0,esk4_0)))|~m1_subset_1(k3_tarski(k2_tarski(esk3_0,esk4_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_104]), c_0_105]), c_0_82]), c_0_85]), c_0_67]), c_0_68])])).
cnf(c_0_111, plain, (m1_subset_1(k3_tarski(k2_tarski(X1,X2)),k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(spm,[status(thm)],[c_0_106, c_0_107])).
cnf(c_0_112, negated_conjecture, (m1_subset_1(esk2_0,k3_tarski(k2_tarski(esk3_0,X1)))), inference(spm,[status(thm)],[c_0_108, c_0_58])).
cnf(c_0_113, negated_conjecture, (~v1_xboole_0(k3_tarski(k2_tarski(esk3_0,X1)))), inference(spm,[status(thm)],[c_0_109, c_0_58])).
cnf(c_0_114, negated_conjecture, (~r2_hidden(esk2_0,k3_tarski(k2_tarski(esk3_0,esk4_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_82]), c_0_85])])).
cnf(c_0_115, negated_conjecture, (r2_hidden(esk2_0,k3_tarski(k2_tarski(esk3_0,X1)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_112]), c_0_113])).
cnf(c_0_116, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_114, c_0_115])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 117
# Proof object clause steps            : 76
# Proof object formula steps           : 41
# Proof object conjectures             : 37
# Proof object clause conjectures      : 34
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 32
# Proof object simplifying inferences  : 68
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 33
# Processed clauses                    : 458
# ...of these trivial                  : 10
# ...subsumed                          : 161
# ...remaining for further processing  : 287
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 17
# Backward-rewritten                   : 61
# Generated clauses                    : 917
# ...of the previous two non-trivial   : 837
# Contextual simplify-reflections      : 5
# Paramodulations                      : 899
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 161
#    Positive orientable unit clauses  : 31
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 13
#    Non-unit-clauses                  : 117
# Current number of unprocessed clauses: 334
# ...number of literals in the above   : 1080
# Current number of archived formulas  : 0
# Current number of archived clauses   : 126
# Clause-clause subsumption calls (NU) : 6947
# Rec. Clause-clause subsumption calls : 3936
# Non-unit clause-clause subsumptions  : 150
# Unit Clause-clause subsumption calls : 441
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 28
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 14515

# -------------------------------------------------
# User time                : 0.023 s
# System time              : 0.000 s
# Total time               : 0.023 s
# Maximum resident set size: 28852 pages
