# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t24_yellow_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1)), file('data/MPTP/proofs/0b07b8c0e45e4331a8d97d73a2935476.E_input', t24_yellow_1)).
fof(t56_setfam_1, axiom, ![X1, X2, X3]:((r1_tarski(k3_tarski(X1),X2)&r2_hidden(X3,X1))=>r1_tarski(X3,X2)), file('data/MPTP/proofs/0b07b8c0e45e4331a8d97d73a2935476.E_input', t56_setfam_1)).
fof(t95_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>r1_tarski(k3_tarski(X1),k3_tarski(X2))), file('data/MPTP/proofs/0b07b8c0e45e4331a8d97d73a2935476.E_input', t95_zfmisc_1)).
fof(d1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1))))))))), file('data/MPTP/proofs/0b07b8c0e45e4331a8d97d73a2935476.E_input', d1_pre_topc)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/0b07b8c0e45e4331a8d97d73a2935476.E_input', d10_xboole_0)).
fof(d12_setfam_1, axiom, ![X1, X2]:(m1_setfam_1(X2,X1)<=>r1_tarski(X1,k3_tarski(X2))), file('data/MPTP/proofs/0b07b8c0e45e4331a8d97d73a2935476.E_input', d12_setfam_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/0b07b8c0e45e4331a8d97d73a2935476.E_input', t3_subset)).
fof(t60_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>(m1_setfam_1(X2,X1)<=>k5_setfam_1(X1,X2)=X1)), file('data/MPTP/proofs/0b07b8c0e45e4331a8d97d73a2935476.E_input', t60_setfam_1)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/0b07b8c0e45e4331a8d97d73a2935476.E_input', t5_subset)).
fof(redefinition_k5_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>k5_setfam_1(X1,X2)=k3_tarski(X2)), file('data/MPTP/proofs/0b07b8c0e45e4331a8d97d73a2935476.E_input', redefinition_k5_setfam_1)).
fof(dt_u1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/0b07b8c0e45e4331a8d97d73a2935476.E_input', dt_u1_pre_topc)).
fof(t14_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>(r2_hidden(k3_tarski(X1),X1)=>k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1))), file('data/MPTP/proofs/0b07b8c0e45e4331a8d97d73a2935476.E_input', t14_yellow_1)).
fof(c_0_12, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1))), inference(assume_negation,[status(cth)],[t24_yellow_1])).
fof(c_0_13, plain, ![X5, X6, X7]:(~r1_tarski(k3_tarski(X5),X6)|~r2_hidden(X7,X5)|r1_tarski(X7,X6)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t56_setfam_1])])).
fof(c_0_14, plain, ![X8, X9]:(~r1_tarski(X8,X9)|r1_tarski(k3_tarski(X8),k3_tarski(X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t95_zfmisc_1])])).
fof(c_0_15, plain, ![X10, X11, X12, X13]:(((r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|~v2_pre_topc(X10)|~l1_pre_topc(X10))&((~m1_subset_1(X11,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|(~r1_tarski(X11,u1_pre_topc(X10))|r2_hidden(k5_setfam_1(u1_struct_0(X10),X11),u1_pre_topc(X10)))|~v2_pre_topc(X10)|~l1_pre_topc(X10))&(~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X10)))|(~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X10)))|(~r2_hidden(X12,u1_pre_topc(X10))|~r2_hidden(X13,u1_pre_topc(X10))|r2_hidden(k9_subset_1(u1_struct_0(X10),X12,X13),u1_pre_topc(X10))))|~v2_pre_topc(X10)|~l1_pre_topc(X10))))&(((m1_subset_1(esk3_1(X10),k1_zfmisc_1(u1_struct_0(X10)))|m1_subset_1(esk2_1(X10),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&((m1_subset_1(esk4_1(X10),k1_zfmisc_1(u1_struct_0(X10)))|m1_subset_1(esk2_1(X10),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&(((r2_hidden(esk3_1(X10),u1_pre_topc(X10))|m1_subset_1(esk2_1(X10),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&(r2_hidden(esk4_1(X10),u1_pre_topc(X10))|m1_subset_1(esk2_1(X10),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10)))&(~r2_hidden(k9_subset_1(u1_struct_0(X10),esk3_1(X10),esk4_1(X10)),u1_pre_topc(X10))|m1_subset_1(esk2_1(X10),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10)))))&(((m1_subset_1(esk3_1(X10),k1_zfmisc_1(u1_struct_0(X10)))|r1_tarski(esk2_1(X10),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&((m1_subset_1(esk4_1(X10),k1_zfmisc_1(u1_struct_0(X10)))|r1_tarski(esk2_1(X10),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&(((r2_hidden(esk3_1(X10),u1_pre_topc(X10))|r1_tarski(esk2_1(X10),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&(r2_hidden(esk4_1(X10),u1_pre_topc(X10))|r1_tarski(esk2_1(X10),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10)))&(~r2_hidden(k9_subset_1(u1_struct_0(X10),esk3_1(X10),esk4_1(X10)),u1_pre_topc(X10))|r1_tarski(esk2_1(X10),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10)))))&((m1_subset_1(esk3_1(X10),k1_zfmisc_1(u1_struct_0(X10)))|~r2_hidden(k5_setfam_1(u1_struct_0(X10),esk2_1(X10)),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&((m1_subset_1(esk4_1(X10),k1_zfmisc_1(u1_struct_0(X10)))|~r2_hidden(k5_setfam_1(u1_struct_0(X10),esk2_1(X10)),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&(((r2_hidden(esk3_1(X10),u1_pre_topc(X10))|~r2_hidden(k5_setfam_1(u1_struct_0(X10),esk2_1(X10)),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&(r2_hidden(esk4_1(X10),u1_pre_topc(X10))|~r2_hidden(k5_setfam_1(u1_struct_0(X10),esk2_1(X10)),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10)))&(~r2_hidden(k9_subset_1(u1_struct_0(X10),esk3_1(X10),esk4_1(X10)),u1_pre_topc(X10))|~r2_hidden(k5_setfam_1(u1_struct_0(X10),esk2_1(X10)),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])).
fof(c_0_16, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&k4_yellow_0(k2_yellow_1(u1_pre_topc(esk1_0)))!=u1_struct_0(esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])).
cnf(c_0_17, plain, (r1_tarski(X3,X2)|~r1_tarski(k3_tarski(X1),X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, plain, (r1_tarski(k3_tarski(X1),k3_tarski(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, plain, ![X17, X18]:(((r1_tarski(X17,X18)|X17!=X18)&(r1_tarski(X18,X17)|X17!=X18))&(~r1_tarski(X17,X18)|~r1_tarski(X18,X17)|X17=X18)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_23, plain, ![X19, X20]:((~m1_setfam_1(X20,X19)|r1_tarski(X19,k3_tarski(X20)))&(~r1_tarski(X19,k3_tarski(X20))|m1_setfam_1(X20,X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_setfam_1])])).
cnf(c_0_24, plain, (r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_25, negated_conjecture, (r2_hidden(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])])).
fof(c_0_26, plain, ![X21, X22]:((~m1_subset_1(X21,k1_zfmisc_1(X22))|r1_tarski(X21,X22))&(~r1_tarski(X21,X22)|m1_subset_1(X21,k1_zfmisc_1(X22)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_27, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_28, plain, ![X23, X24]:((~m1_setfam_1(X24,X23)|k5_setfam_1(X23,X24)=X23|~m1_subset_1(X24,k1_zfmisc_1(k1_zfmisc_1(X23))))&(k5_setfam_1(X23,X24)!=X23|m1_setfam_1(X24,X23)|~m1_subset_1(X24,k1_zfmisc_1(k1_zfmisc_1(X23))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t60_setfam_1])])])).
cnf(c_0_29, plain, (m1_setfam_1(X2,X1)|~r1_tarski(X1,k3_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30, negated_conjecture, (r1_tarski(u1_struct_0(esk1_0),k3_tarski(X1))|~r1_tarski(u1_pre_topc(esk1_0),X1)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
fof(c_0_31, plain, ![X25, X26, X27]:(~r2_hidden(X25,X26)|(~m1_subset_1(X26,k1_zfmisc_1(X27))|~v1_xboole_0(X27))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
cnf(c_0_32, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_27])).
fof(c_0_34, plain, ![X28, X29]:(~m1_subset_1(X29,k1_zfmisc_1(k1_zfmisc_1(X28)))|k5_setfam_1(X28,X29)=k3_tarski(X29)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k5_setfam_1])])).
fof(c_0_35, plain, ![X30]:(~l1_pre_topc(X30)|m1_subset_1(u1_pre_topc(X30),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X30))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_pre_topc])])).
cnf(c_0_36, plain, (k5_setfam_1(X2,X1)=X2|~m1_setfam_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, negated_conjecture, (m1_setfam_1(X1,u1_struct_0(esk1_0))|~r1_tarski(u1_pre_topc(esk1_0),X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
fof(c_0_38, plain, ![X31]:(v1_xboole_0(X31)|(~r2_hidden(k3_tarski(X31),X31)|k4_yellow_0(k2_yellow_1(X31))=k3_tarski(X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_yellow_1])])])).
cnf(c_0_39, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_40, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_41, plain, (k5_setfam_1(X2,X1)=k3_tarski(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_43, negated_conjecture, (k5_setfam_1(u1_struct_0(esk1_0),X1)=u1_struct_0(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~r1_tarski(u1_pre_topc(esk1_0),X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_44, plain, (v1_xboole_0(X1)|k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_45, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_46, plain, (k5_setfam_1(u1_struct_0(X1),u1_pre_topc(X1))=k3_tarski(u1_pre_topc(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_47, negated_conjecture, (k5_setfam_1(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_42]), c_0_33]), c_0_21])])).
cnf(c_0_48, plain, (k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(csr,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_49, negated_conjecture, (k3_tarski(u1_pre_topc(esk1_0))=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_21])])).
cnf(c_0_50, negated_conjecture, (k4_yellow_0(k2_yellow_1(u1_pre_topc(esk1_0)))!=u1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_51, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_25])]), c_0_50]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 52
# Proof object clause steps            : 27
# Proof object formula steps           : 25
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 37
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 37
# Processed clauses                    : 112
# ...of these trivial                  : 0
# ...subsumed                          : 6
# ...remaining for further processing  : 106
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 129
# ...of the previous two non-trivial   : 100
# Contextual simplify-reflections      : 1
# Paramodulations                      : 127
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 67
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 54
# Current number of unprocessed clauses: 58
# ...number of literals in the above   : 215
# Current number of archived formulas  : 0
# Current number of archived clauses   : 37
# Clause-clause subsumption calls (NU) : 751
# Rec. Clause-clause subsumption calls : 186
# Non-unit clause-clause subsumptions  : 8
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5555

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28728 pages
