# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/0b29ab2f0a9b4a80b5d33ba3c487bf02.E_input', t4_xboole_0)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/0b29ab2f0a9b4a80b5d33ba3c487bf02.E_input', t48_xboole_1)).
fof(t8_boole, axiom, ![X1, X2]:~((v1_xboole_0(X1)&(~(X1=X2)&v1_xboole_0(X2)))), file('data/MPTP/proofs/0b29ab2f0a9b4a80b5d33ba3c487bf02.E_input', t8_boole)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/0b29ab2f0a9b4a80b5d33ba3c487bf02.E_input', d1_xboole_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/0b29ab2f0a9b4a80b5d33ba3c487bf02.E_input', fc1_xboole_0)).
fof(t72_xboole_1, conjecture, ![X1, X2, X3, X4]:((k2_xboole_0(X1,X2)=k2_xboole_0(X3,X4)&(r1_xboole_0(X3,X1)&r1_xboole_0(X4,X2)))=>X3=X2), file('data/MPTP/proofs/0b29ab2f0a9b4a80b5d33ba3c487bf02.E_input', t72_xboole_1)).
fof(t51_xboole_1, axiom, ![X1, X2]:k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/0b29ab2f0a9b4a80b5d33ba3c487bf02.E_input', t51_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/0b29ab2f0a9b4a80b5d33ba3c487bf02.E_input', commutativity_k2_xboole_0)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/0b29ab2f0a9b4a80b5d33ba3c487bf02.E_input', t39_xboole_1)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/0b29ab2f0a9b4a80b5d33ba3c487bf02.E_input', t1_boole)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/0b29ab2f0a9b4a80b5d33ba3c487bf02.E_input', t2_boole)).
fof(t6_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k2_xboole_0(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/0b29ab2f0a9b4a80b5d33ba3c487bf02.E_input', t6_xboole_1)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/0b29ab2f0a9b4a80b5d33ba3c487bf02.E_input', t41_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/0b29ab2f0a9b4a80b5d33ba3c487bf02.E_input', t3_boole)).
fof(t42_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k4_xboole_0(X1,X3),k4_xboole_0(X2,X3)), file('data/MPTP/proofs/0b29ab2f0a9b4a80b5d33ba3c487bf02.E_input', t42_xboole_1)).
fof(t4_xboole_1, axiom, ![X1, X2, X3]:k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/0b29ab2f0a9b4a80b5d33ba3c487bf02.E_input', t4_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/0b29ab2f0a9b4a80b5d33ba3c487bf02.E_input', commutativity_k3_xboole_0)).
fof(c_0_17, plain, ![X31, X32, X34, X35, X36]:((r1_xboole_0(X31,X32)|r2_hidden(esk6_2(X31,X32),k3_xboole_0(X31,X32)))&(~r2_hidden(X36,k3_xboole_0(X34,X35))|~r1_xboole_0(X34,X35))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_18, plain, ![X28, X29]:k4_xboole_0(X28,k4_xboole_0(X28,X29))=k3_xboole_0(X28,X29), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_19, plain, ![X43, X44]:(~v1_xboole_0(X43)|(X43=X44|~v1_xboole_0(X44))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_boole])])])).
cnf(c_0_20, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_22, plain, ![X12, X13, X14]:((~v1_xboole_0(X12)|~r2_hidden(X13,X12))&(r2_hidden(esk5_1(X14),X14)|v1_xboole_0(X14))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_23, plain, (X1=X2|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_25, plain, (~r1_xboole_0(X2,X3)|~r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_26, plain, (r2_hidden(esk5_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_27, negated_conjecture, ~(![X1, X2, X3, X4]:((k2_xboole_0(X1,X2)=k2_xboole_0(X3,X4)&(r1_xboole_0(X3,X1)&r1_xboole_0(X4,X2)))=>X3=X2)), inference(assume_negation,[status(cth)],[t72_xboole_1])).
cnf(c_0_28, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_29, plain, (v1_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_30, negated_conjecture, ((k2_xboole_0(esk1_0,esk2_0)=k2_xboole_0(esk3_0,esk4_0)&(r1_xboole_0(esk3_0,esk1_0)&r1_xboole_0(esk4_0,esk2_0)))&esk3_0!=esk2_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])])).
fof(c_0_31, plain, ![X26, X27]:k2_xboole_0(k3_xboole_0(X26,X27),k4_xboole_0(X26,X27))=X26, inference(variable_rename,[status(thm)],[t51_xboole_1])).
fof(c_0_32, plain, ![X17, X18]:k2_xboole_0(X17,X18)=k2_xboole_0(X18,X17), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_33, plain, ![X24, X25]:k2_xboole_0(X24,k4_xboole_0(X25,X24))=k2_xboole_0(X24,X25), inference(variable_rename,[status(thm)],[t39_xboole_1])).
cnf(c_0_34, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35, negated_conjecture, (r1_xboole_0(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_36, plain, ![X16]:k2_xboole_0(X16,k1_xboole_0)=X16, inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_37, plain, (k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_38, plain, ![X30]:k3_xboole_0(X30,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_39, plain, ![X22, X23]:k2_xboole_0(X22,k2_xboole_0(X22,X23))=k2_xboole_0(X22,X23), inference(variable_rename,[status(thm)],[t6_xboole_1])).
fof(c_0_40, plain, ![X9, X10, X11]:k4_xboole_0(k4_xboole_0(X9,X10),X11)=k4_xboole_0(X9,k2_xboole_0(X10,X11)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
cnf(c_0_41, negated_conjecture, (k2_xboole_0(esk1_0,esk2_0)=k2_xboole_0(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_42, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_44, negated_conjecture, (k4_xboole_0(esk4_0,k4_xboole_0(esk4_0,esk2_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_45, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_46, plain, (k2_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k4_xboole_0(X1,X2))=X1), inference(rw,[status(thm)],[c_0_37, c_0_21])).
cnf(c_0_47, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_48, plain, ![X21]:k4_xboole_0(X21,k1_xboole_0)=X21, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_49, negated_conjecture, (r1_xboole_0(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_50, plain, (k2_xboole_0(X1,k2_xboole_0(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_52, negated_conjecture, (k2_xboole_0(esk2_0,esk1_0)=k2_xboole_0(esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_53, negated_conjecture, (k2_xboole_0(esk4_0,k4_xboole_0(esk4_0,esk2_0))=k4_xboole_0(esk4_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_42])).
cnf(c_0_54, plain, (k2_xboole_0(X1,k4_xboole_0(X1,X2))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_42]), c_0_43]), c_0_42])).
cnf(c_0_55, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k1_xboole_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_47, c_0_21])).
cnf(c_0_56, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_57, plain, ![X40, X41, X42]:k4_xboole_0(k2_xboole_0(X40,X41),X42)=k2_xboole_0(k4_xboole_0(X40,X42),k4_xboole_0(X41,X42)), inference(variable_rename,[status(thm)],[t42_xboole_1])).
cnf(c_0_58, negated_conjecture, (k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,esk1_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_34, c_0_49])).
cnf(c_0_59, plain, (k2_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_50, c_0_45])).
cnf(c_0_60, negated_conjecture, (k4_xboole_0(k4_xboole_0(X1,esk2_0),esk1_0)=k4_xboole_0(k4_xboole_0(X1,esk3_0),esk4_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_51])).
cnf(c_0_61, negated_conjecture, (k4_xboole_0(esk4_0,esk2_0)=esk4_0), inference(rw,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_62, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_63, plain, (k4_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k4_xboole_0(X1,X3),k4_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_64, plain, (k2_xboole_0(k1_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_45, c_0_42])).
cnf(c_0_65, negated_conjecture, (k2_xboole_0(esk3_0,k4_xboole_0(esk3_0,esk1_0))=k4_xboole_0(esk3_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_58]), c_0_45]), c_0_42])).
cnf(c_0_66, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_51, c_0_59])).
cnf(c_0_67, negated_conjecture, (k4_xboole_0(k4_xboole_0(esk4_0,esk3_0),esk4_0)=k4_xboole_0(esk4_0,esk1_0)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_68, plain, (k4_xboole_0(k4_xboole_0(k2_xboole_0(X1,X2),X1),X2)=k1_xboole_0), inference(spm,[status(thm)],[c_0_62, c_0_51])).
cnf(c_0_69, plain, (k4_xboole_0(k2_xboole_0(X1,X2),X1)=k4_xboole_0(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_62]), c_0_64])).
fof(c_0_70, plain, ![X37, X38, X39]:k2_xboole_0(k2_xboole_0(X37,X38),X39)=k2_xboole_0(X37,k2_xboole_0(X38,X39)), inference(variable_rename,[status(thm)],[t4_xboole_1])).
cnf(c_0_71, negated_conjecture, (k4_xboole_0(esk3_0,esk1_0)=esk3_0), inference(rw,[status(thm)],[c_0_65, c_0_54])).
cnf(c_0_72, negated_conjecture, (k4_xboole_0(k4_xboole_0(esk4_0,esk1_0),esk4_0)=k4_xboole_0(esk4_0,esk1_0)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_73, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X1)=k1_xboole_0), inference(rw,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_74, plain, (k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_62]), c_0_45])).
cnf(c_0_75, plain, (k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_76, negated_conjecture, (k2_xboole_0(esk3_0,k4_xboole_0(X1,esk1_0))=k4_xboole_0(k2_xboole_0(esk3_0,X1),esk1_0)), inference(spm,[status(thm)],[c_0_63, c_0_71])).
cnf(c_0_77, negated_conjecture, (k4_xboole_0(esk4_0,esk1_0)=k1_xboole_0), inference(rw,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_78, negated_conjecture, (k4_xboole_0(k2_xboole_0(esk3_0,esk4_0),esk1_0)=k4_xboole_0(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_74, c_0_52])).
cnf(c_0_79, negated_conjecture, (k2_xboole_0(esk2_0,k2_xboole_0(esk1_0,X1))=k2_xboole_0(esk3_0,k2_xboole_0(esk4_0,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_52]), c_0_75])).
fof(c_0_80, plain, ![X19, X20]:k3_xboole_0(X19,X20)=k3_xboole_0(X20,X19), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_81, negated_conjecture, (k4_xboole_0(esk2_0,esk1_0)=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_45]), c_0_78])).
cnf(c_0_82, negated_conjecture, (k2_xboole_0(esk3_0,k2_xboole_0(esk1_0,esk4_0))=k2_xboole_0(esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_59]), c_0_52]), c_0_42])).
cnf(c_0_83, negated_conjecture, (k2_xboole_0(esk1_0,esk4_0)=esk1_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_77]), c_0_45])).
cnf(c_0_84, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_85, negated_conjecture, (k2_xboole_0(esk4_0,k4_xboole_0(X1,esk2_0))=k4_xboole_0(k2_xboole_0(esk4_0,X1),esk2_0)), inference(spm,[status(thm)],[c_0_63, c_0_61])).
cnf(c_0_86, negated_conjecture, (k4_xboole_0(esk3_0,esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_73, c_0_81])).
cnf(c_0_87, negated_conjecture, (k2_xboole_0(esk3_0,esk4_0)=k2_xboole_0(esk3_0,esk1_0)), inference(rw,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_88, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_84, c_0_21]), c_0_21])).
cnf(c_0_89, negated_conjecture, (k4_xboole_0(k2_xboole_0(esk3_0,esk1_0),esk2_0)=esk4_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_45]), c_0_42]), c_0_87])).
cnf(c_0_90, negated_conjecture, (k2_xboole_0(esk2_0,k2_xboole_0(X1,esk1_0))=k2_xboole_0(esk3_0,k2_xboole_0(esk4_0,X1))), inference(spm,[status(thm)],[c_0_79, c_0_42])).
cnf(c_0_91, negated_conjecture, (k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk4_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_61]), c_0_62])).
cnf(c_0_92, negated_conjecture, (k2_xboole_0(esk2_0,esk4_0)=k2_xboole_0(esk3_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_89]), c_0_90]), c_0_42]), c_0_87]), c_0_50])).
cnf(c_0_93, negated_conjecture, (k4_xboole_0(esk2_0,esk4_0)=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_91]), c_0_45]), c_0_42]), c_0_54])).
cnf(c_0_94, negated_conjecture, (k4_xboole_0(k2_xboole_0(esk3_0,esk1_0),esk4_0)=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_92]), c_0_93])).
cnf(c_0_95, negated_conjecture, (k4_xboole_0(esk3_0,esk4_0)=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_87]), c_0_94])).
cnf(c_0_96, negated_conjecture, (k2_xboole_0(esk3_0,esk2_0)=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_81]), c_0_42])).
cnf(c_0_97, negated_conjecture, (esk3_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_98, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_95]), c_0_96]), c_0_97]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 99
# Proof object clause steps            : 65
# Proof object formula steps           : 34
# Proof object conjectures             : 36
# Proof object clause conjectures      : 33
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 33
# Proof object simplifying inferences  : 44
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 21
# Processed clauses                    : 837
# ...of these trivial                  : 46
# ...subsumed                          : 558
# ...remaining for further processing  : 233
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 70
# Generated clauses                    : 3724
# ...of the previous two non-trivial   : 3046
# Contextual simplify-reflections      : 0
# Paramodulations                      : 3724
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 140
#    Positive orientable unit clauses  : 72
#    Positive unorientable unit clauses: 5
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 61
# Current number of unprocessed clauses: 2221
# ...number of literals in the above   : 3066
# Current number of archived formulas  : 0
# Current number of archived clauses   : 94
# Clause-clause subsumption calls (NU) : 4555
# Rec. Clause-clause subsumption calls : 3001
# Non-unit clause-clause subsumptions  : 490
# Unit Clause-clause subsumption calls : 26
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 227
# BW rewrite match successes           : 140
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 39224

# -------------------------------------------------
# User time                : 0.027 s
# System time              : 0.003 s
# Total time               : 0.030 s
# Maximum resident set size: 28748 pages
