# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.011 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', d1_zfmisc_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', d3_tarski)).
fof(t3_xboole_1, axiom, ![X1]:(r1_tarski(X1,k1_xboole_0)=>X1=k1_xboole_0), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', t3_xboole_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', t2_xboole_1)).
fof(fraenkel_a_2_6_lattice3, axiom, ![X1, X2, X3]:((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))=>(r2_hidden(X1,a_2_6_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r3_lattices(X2,X5,X4)&r2_hidden(X5,X3))))))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', fraenkel_a_2_6_lattice3)).
fof(t34_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(X2=k16_lattice3(X1,X3)<=>(r3_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_lattice3(X1,X4,X3)=>r3_lattices(X1,X4,X2))))))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', t34_lattice3)).
fof(dt_k16_lattice3, axiom, ![X1, X2]:((~(v2_struct_0(X1))&l3_lattices(X1))=>m1_subset_1(k16_lattice3(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', dt_k16_lattice3)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', t7_ordinal1)).
fof(d16_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r3_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X2,X4)))))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', d16_lattice3)).
fof(fraenkel_a_2_5_lattice3, axiom, ![X1, X2, X3]:((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))=>(r2_hidden(X1,a_2_5_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r3_lattices(X2,X4,X5)&r2_hidden(X5,X3))))))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', fraenkel_a_2_5_lattice3)).
fof(t50_lattice3, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))&k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', t50_lattice3)).
fof(t37_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:k15_lattice3(X1,X2)=k16_lattice3(X1,a_2_2_lattice3(X1,X2))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', t37_lattice3)).
fof(t47_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(k15_lattice3(X1,X2)=k15_lattice3(X1,a_2_5_lattice3(X1,X2))&k16_lattice3(X1,X2)=k16_lattice3(X1,a_2_6_lattice3(X1,X2)))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', t47_lattice3)).
fof(t48_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X2)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>~((r3_lattices(X1,X4,X5)&r2_hidden(X5,X3)))))))=>r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3)))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', t48_lattice3)).
fof(d21_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X3=k15_lattice3(X1,X2)<=>(r4_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r4_lattice3(X1,X4,X2)=>r1_lattices(X1,X3,X4)))))))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', d21_lattice3)).
fof(dt_k15_lattice3, axiom, ![X1, X2]:((~(v2_struct_0(X1))&l3_lattices(X1))=>m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', dt_k15_lattice3)).
fof(t42_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((r2_hidden(X2,X3)&r3_lattice3(X1,X2,X3))=>k16_lattice3(X1,X3)=X2))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', t42_lattice3)).
fof(fraenkel_a_2_2_lattice3, axiom, ![X1, X2, X3]:((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))=>(r2_hidden(X1,a_2_2_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&r4_lattice3(X2,X4,X3))))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', fraenkel_a_2_2_lattice3)).
fof(t43_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k15_lattice3(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k16_lattice3(X1,k6_domain_1(u1_struct_0(X1),X2))=X2))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', t43_lattice3)).
fof(d16_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v13_lattices(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(X2=k5_lattices(X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_lattices(X1,X2,X3)=X2&k2_lattices(X1,X3,X2)=X2)))))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', d16_lattices)).
fof(dt_k5_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>m1_subset_1(k5_lattices(X1),u1_struct_0(X1))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', dt_k5_lattices)).
fof(t21_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattices(X1,X2,X3)<=>k2_lattices(X1,X2,X3)=X2)))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', t21_lattices)).
fof(d6_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v6_lattices(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2))))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', d6_lattices)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', cc1_lattices)).
fof(dt_l3_lattices, axiom, ![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', dt_l3_lattices)).
fof(d13_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v13_lattices(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_lattices(X1,X2,X3)=X2&k2_lattices(X1,X3,X2)=X2))))), file('data/MPTP/proofs/0b47f23365e24785819dac1c32754035.E_input', d13_lattices)).
fof(c_0_26, plain, ![X6, X7, X8, X9, X10, X11]:(((~r2_hidden(X8,X7)|r1_tarski(X8,X6)|X7!=k1_zfmisc_1(X6))&(~r1_tarski(X9,X6)|r2_hidden(X9,X7)|X7!=k1_zfmisc_1(X6)))&((~r2_hidden(esk1_2(X10,X11),X11)|~r1_tarski(esk1_2(X10,X11),X10)|X11=k1_zfmisc_1(X10))&(r2_hidden(esk1_2(X10,X11),X11)|r1_tarski(esk1_2(X10,X11),X10)|X11=k1_zfmisc_1(X10)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
cnf(c_0_27, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_28, plain, ![X13, X14, X15, X16, X17]:((~r1_tarski(X13,X14)|(~r2_hidden(X15,X13)|r2_hidden(X15,X14)))&((r2_hidden(esk2_2(X16,X17),X16)|r1_tarski(X16,X17))&(~r2_hidden(esk2_2(X16,X17),X17)|r1_tarski(X16,X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_29, plain, (r2_hidden(X1,X3)|~r1_tarski(X1,X2)|X3!=k1_zfmisc_1(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_30, plain, ![X19]:(~r1_tarski(X19,k1_xboole_0)|X19=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_xboole_1])])).
cnf(c_0_31, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_32, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33, plain, (r1_tarski(X1,X2)|~r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34, plain, (r2_hidden(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(er,[status(thm)],[c_0_29])).
cnf(c_0_35, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_36, plain, (r1_tarski(esk2_2(k1_zfmisc_1(X1),X2),X1)|r1_tarski(k1_zfmisc_1(X1),X2)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
fof(c_0_37, plain, ![X20]:r1_tarski(k1_xboole_0,X20), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_38, plain, (r1_tarski(X1,k1_zfmisc_1(X2))|~r1_tarski(esk2_2(X1,k1_zfmisc_1(X2)),X2)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_39, plain, (esk2_2(k1_zfmisc_1(k1_xboole_0),X1)=k1_xboole_0|r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_40, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_41, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_42, plain, (r1_tarski(k1_zfmisc_1(k1_xboole_0),k1_zfmisc_1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])])).
fof(c_0_43, plain, ![X21, X22, X23, X26, X27, X28]:(((m1_subset_1(esk3_3(X21,X22,X23),u1_struct_0(X22))|~r2_hidden(X21,a_2_6_lattice3(X22,X23))|(v2_struct_0(X22)|(~v10_lattices(X22)|(~v4_lattice3(X22)|~l3_lattices(X22)))))&((X21=esk3_3(X21,X22,X23)|~r2_hidden(X21,a_2_6_lattice3(X22,X23))|(v2_struct_0(X22)|(~v10_lattices(X22)|(~v4_lattice3(X22)|~l3_lattices(X22)))))&((m1_subset_1(esk4_3(X21,X22,X23),u1_struct_0(X22))|~r2_hidden(X21,a_2_6_lattice3(X22,X23))|(v2_struct_0(X22)|(~v10_lattices(X22)|(~v4_lattice3(X22)|~l3_lattices(X22)))))&((r3_lattices(X22,esk4_3(X21,X22,X23),esk3_3(X21,X22,X23))|~r2_hidden(X21,a_2_6_lattice3(X22,X23))|(v2_struct_0(X22)|(~v10_lattices(X22)|(~v4_lattice3(X22)|~l3_lattices(X22)))))&(r2_hidden(esk4_3(X21,X22,X23),X23)|~r2_hidden(X21,a_2_6_lattice3(X22,X23))|(v2_struct_0(X22)|(~v10_lattices(X22)|(~v4_lattice3(X22)|~l3_lattices(X22)))))))))&(~m1_subset_1(X27,u1_struct_0(X22))|(X21!=X27|(~m1_subset_1(X28,u1_struct_0(X22))|(~r3_lattices(X22,X28,X27)|~r2_hidden(X28,X26))))|r2_hidden(X21,a_2_6_lattice3(X22,X26))|(v2_struct_0(X22)|(~v10_lattices(X22)|(~v4_lattice3(X22)|~l3_lattices(X22)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_6_lattice3])])])])])])])).
fof(c_0_44, plain, ![X29, X30, X31, X32, X33]:(((r3_lattice3(X29,X30,X31)|X30!=k16_lattice3(X29,X31)|~m1_subset_1(X30,u1_struct_0(X29))|(v2_struct_0(X29)|(~v10_lattices(X29)|(~v4_lattice3(X29)|~l3_lattices(X29)))))&(~m1_subset_1(X32,u1_struct_0(X29))|(~r3_lattice3(X29,X32,X31)|r3_lattices(X29,X32,X30))|X30!=k16_lattice3(X29,X31)|~m1_subset_1(X30,u1_struct_0(X29))|(v2_struct_0(X29)|(~v10_lattices(X29)|(~v4_lattice3(X29)|~l3_lattices(X29))))))&((m1_subset_1(esk5_3(X29,X30,X33),u1_struct_0(X29))|~r3_lattice3(X29,X30,X33)|X30=k16_lattice3(X29,X33)|~m1_subset_1(X30,u1_struct_0(X29))|(v2_struct_0(X29)|(~v10_lattices(X29)|(~v4_lattice3(X29)|~l3_lattices(X29)))))&((r3_lattice3(X29,esk5_3(X29,X30,X33),X33)|~r3_lattice3(X29,X30,X33)|X30=k16_lattice3(X29,X33)|~m1_subset_1(X30,u1_struct_0(X29))|(v2_struct_0(X29)|(~v10_lattices(X29)|(~v4_lattice3(X29)|~l3_lattices(X29)))))&(~r3_lattices(X29,esk5_3(X29,X30,X33),X30)|~r3_lattice3(X29,X30,X33)|X30=k16_lattice3(X29,X33)|~m1_subset_1(X30,u1_struct_0(X29))|(v2_struct_0(X29)|(~v10_lattices(X29)|(~v4_lattice3(X29)|~l3_lattices(X29)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_lattice3])])])])])])])).
fof(c_0_45, plain, ![X35, X36]:(v2_struct_0(X35)|~l3_lattices(X35)|m1_subset_1(k16_lattice3(X35,X36),u1_struct_0(X35))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k16_lattice3])])])).
fof(c_0_46, plain, ![X37, X38]:(~r2_hidden(X37,X38)|~r1_tarski(X38,X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_47, plain, (r2_hidden(X1,k1_zfmisc_1(X2))|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_48, plain, (r2_hidden(X3,a_2_6_lattice3(X2,X5))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~m1_subset_1(X4,u1_struct_0(X2))|~r3_lattices(X2,X4,X1)|~r2_hidden(X4,X5)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49, plain, (r3_lattices(X2,X1,X4)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r3_lattice3(X2,X1,X3)|X4!=k16_lattice3(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_50, plain, (v2_struct_0(X1)|m1_subset_1(k16_lattice3(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_51, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_52, plain, ![X39, X40, X41, X42, X43]:((~r3_lattice3(X39,X40,X41)|(~m1_subset_1(X42,u1_struct_0(X39))|(~r2_hidden(X42,X41)|r1_lattices(X39,X40,X42)))|~m1_subset_1(X40,u1_struct_0(X39))|(v2_struct_0(X39)|~l3_lattices(X39)))&((m1_subset_1(esk6_3(X39,X40,X43),u1_struct_0(X39))|r3_lattice3(X39,X40,X43)|~m1_subset_1(X40,u1_struct_0(X39))|(v2_struct_0(X39)|~l3_lattices(X39)))&((r2_hidden(esk6_3(X39,X40,X43),X43)|r3_lattice3(X39,X40,X43)|~m1_subset_1(X40,u1_struct_0(X39))|(v2_struct_0(X39)|~l3_lattices(X39)))&(~r1_lattices(X39,X40,esk6_3(X39,X40,X43))|r3_lattice3(X39,X40,X43)|~m1_subset_1(X40,u1_struct_0(X39))|(v2_struct_0(X39)|~l3_lattices(X39)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattice3])])])])])])])).
cnf(c_0_53, plain, (r2_hidden(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_47, c_0_34])).
cnf(c_0_54, plain, (v2_struct_0(X1)|r2_hidden(X2,a_2_6_lattice3(X1,X3))|~r3_lattices(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(X4,X3)), inference(er,[status(thm)],[c_0_48])).
cnf(c_0_55, plain, (r3_lattices(X1,X2,k16_lattice3(X1,X3))|v2_struct_0(X1)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_49]), c_0_50])).
cnf(c_0_56, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_51, c_0_40])).
cnf(c_0_57, plain, (r2_hidden(esk6_3(X1,X2,X3),X3)|r3_lattice3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_58, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_31, c_0_53])).
cnf(c_0_59, plain, (v2_struct_0(X1)|r2_hidden(k16_lattice3(X1,X2),a_2_6_lattice3(X1,X3))|~r3_lattice3(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(X4,X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_50])).
cnf(c_0_60, plain, (r3_lattice3(X1,X2,k1_xboole_0)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
fof(c_0_61, plain, ![X45, X46, X47, X50, X51, X52]:(((m1_subset_1(esk7_3(X45,X46,X47),u1_struct_0(X46))|~r2_hidden(X45,a_2_5_lattice3(X46,X47))|(v2_struct_0(X46)|(~v10_lattices(X46)|(~v4_lattice3(X46)|~l3_lattices(X46)))))&((X45=esk7_3(X45,X46,X47)|~r2_hidden(X45,a_2_5_lattice3(X46,X47))|(v2_struct_0(X46)|(~v10_lattices(X46)|(~v4_lattice3(X46)|~l3_lattices(X46)))))&((m1_subset_1(esk8_3(X45,X46,X47),u1_struct_0(X46))|~r2_hidden(X45,a_2_5_lattice3(X46,X47))|(v2_struct_0(X46)|(~v10_lattices(X46)|(~v4_lattice3(X46)|~l3_lattices(X46)))))&((r3_lattices(X46,esk7_3(X45,X46,X47),esk8_3(X45,X46,X47))|~r2_hidden(X45,a_2_5_lattice3(X46,X47))|(v2_struct_0(X46)|(~v10_lattices(X46)|(~v4_lattice3(X46)|~l3_lattices(X46)))))&(r2_hidden(esk8_3(X45,X46,X47),X47)|~r2_hidden(X45,a_2_5_lattice3(X46,X47))|(v2_struct_0(X46)|(~v10_lattices(X46)|(~v4_lattice3(X46)|~l3_lattices(X46)))))))))&(~m1_subset_1(X51,u1_struct_0(X46))|(X45!=X51|(~m1_subset_1(X52,u1_struct_0(X46))|(~r3_lattices(X46,X51,X52)|~r2_hidden(X52,X50))))|r2_hidden(X45,a_2_5_lattice3(X46,X50))|(v2_struct_0(X46)|(~v10_lattices(X46)|(~v4_lattice3(X46)|~l3_lattices(X46)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_5_lattice3])])])])])])])).
cnf(c_0_62, plain, (r1_tarski(esk2_2(k1_zfmisc_1(k1_xboole_0),X1),X2)|r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_58, c_0_36])).
cnf(c_0_63, plain, (r2_hidden(esk1_2(X1,X2),X2)|r1_tarski(esk1_2(X1,X2),X1)|X2=k1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_64, plain, (v2_struct_0(X1)|r2_hidden(k16_lattice3(X1,k1_xboole_0),a_2_6_lattice3(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_65, plain, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_5_lattice3(X2,X3))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_66, plain, (r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)|~r2_hidden(X2,esk2_2(k1_zfmisc_1(k1_xboole_0),X1))), inference(spm,[status(thm)],[c_0_51, c_0_62])).
cnf(c_0_67, plain, (X1=k1_zfmisc_1(k1_xboole_0)|r1_tarski(esk1_2(k1_xboole_0,X1),X2)|r2_hidden(esk1_2(k1_xboole_0,X1),X1)), inference(spm,[status(thm)],[c_0_58, c_0_63])).
cnf(c_0_68, plain, (~r2_hidden(k1_zfmisc_1(X1),k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_51, c_0_42])).
cnf(c_0_69, plain, (v2_struct_0(X1)|r2_hidden(k16_lattice3(X1,k1_xboole_0),a_2_6_lattice3(X1,X2))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(esk8_3(X3,X1,X4),X2)|~r2_hidden(X3,a_2_5_lattice3(X1,X4))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_70, plain, (r2_hidden(esk8_3(X1,X2,X3),X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_5_lattice3(X2,X3))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_61])).
fof(c_0_71, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))&k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0)))), inference(assume_negation,[status(cth)],[t50_lattice3])).
cnf(c_0_72, plain, (esk2_2(k1_zfmisc_1(k1_xboole_0),X1)=k1_zfmisc_1(k1_xboole_0)|r1_tarski(esk1_2(k1_xboole_0,esk2_2(k1_zfmisc_1(k1_xboole_0),X1)),X2)|r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_73, plain, (~r1_tarski(k1_zfmisc_1(X1),k1_xboole_0)), inference(spm,[status(thm)],[c_0_68, c_0_34])).
cnf(c_0_74, plain, (r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_32, c_0_39])).
cnf(c_0_75, plain, (v2_struct_0(X1)|r2_hidden(k16_lattice3(X1,k1_xboole_0),a_2_6_lattice3(X1,X2))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(X3,a_2_5_lattice3(X1,X2))), inference(spm,[status(thm)],[c_0_69, c_0_70])).
fof(c_0_76, negated_conjecture, ((~v2_struct_0(esk9_0)&(v10_lattices(esk9_0)&(v4_lattice3(esk9_0)&l3_lattices(esk9_0))))&(v2_struct_0(esk9_0)|(~v10_lattices(esk9_0)|(~v13_lattices(esk9_0)|~l3_lattices(esk9_0)))|k5_lattices(esk9_0)!=k15_lattice3(esk9_0,k1_xboole_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_71])])])])).
cnf(c_0_77, plain, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|r1_tarski(esk1_2(k1_xboole_0,k1_xboole_0),X1)|r1_tarski(k1_zfmisc_1(k1_xboole_0),X2)), inference(spm,[status(thm)],[c_0_72, c_0_39])).
cnf(c_0_78, plain, (r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_79, plain, (r2_hidden(esk4_3(X1,X2,X3),X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_6_lattice3(X2,X3))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_80, plain, (v2_struct_0(X1)|r1_tarski(a_2_5_lattice3(X1,X2),X3)|r2_hidden(k16_lattice3(X1,k1_xboole_0),a_2_6_lattice3(X1,X2))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_75, c_0_32])).
cnf(c_0_81, negated_conjecture, (l3_lattices(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_82, negated_conjecture, (v4_lattice3(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_83, negated_conjecture, (v10_lattices(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_84, negated_conjecture, (~v2_struct_0(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_85, plain, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|r1_tarski(esk1_2(k1_xboole_0,k1_xboole_0),X1)|~r2_hidden(X2,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_51, c_0_77])).
cnf(c_0_86, plain, (r2_hidden(k1_xboole_0,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_47, c_0_78])).
cnf(c_0_87, plain, (v2_struct_0(X1)|r1_tarski(k1_zfmisc_1(k1_xboole_0),X2)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(X3,a_2_6_lattice3(X1,esk2_2(k1_zfmisc_1(k1_xboole_0),X2)))), inference(spm,[status(thm)],[c_0_66, c_0_79])).
cnf(c_0_88, negated_conjecture, (r1_tarski(a_2_5_lattice3(esk9_0,X1),X2)|r2_hidden(k16_lattice3(esk9_0,k1_xboole_0),a_2_6_lattice3(esk9_0,X1))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_82]), c_0_83])]), c_0_84])).
cnf(c_0_89, plain, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|r1_tarski(esk1_2(k1_xboole_0,k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_90, plain, (r3_lattice3(X1,X2,X3)|v2_struct_0(X1)|X2!=k16_lattice3(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_91, plain, ![X54, X55]:(v2_struct_0(X54)|(~v10_lattices(X54)|(~v4_lattice3(X54)|~l3_lattices(X54)))|k15_lattice3(X54,X55)=k16_lattice3(X54,a_2_2_lattice3(X54,X55))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t37_lattice3])])])])).
fof(c_0_92, plain, ![X56, X57]:((k15_lattice3(X56,X57)=k15_lattice3(X56,a_2_5_lattice3(X56,X57))|(v2_struct_0(X56)|(~v10_lattices(X56)|(~v4_lattice3(X56)|~l3_lattices(X56)))))&(k16_lattice3(X56,X57)=k16_lattice3(X56,a_2_6_lattice3(X56,X57))|(v2_struct_0(X56)|(~v10_lattices(X56)|(~v4_lattice3(X56)|~l3_lattices(X56)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t47_lattice3])])])])])).
cnf(c_0_93, negated_conjecture, (r1_tarski(a_2_5_lattice3(esk9_0,esk2_2(k1_zfmisc_1(k1_xboole_0),X1)),X2)|r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_81]), c_0_82]), c_0_83])]), c_0_84])).
cnf(c_0_94, plain, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|~r2_hidden(X1,esk1_2(k1_xboole_0,k1_xboole_0))), inference(spm,[status(thm)],[c_0_51, c_0_89])).
cnf(c_0_95, plain, (r3_lattice3(X1,k16_lattice3(X1,X2),X2)|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_90]), c_0_50])).
cnf(c_0_96, plain, (v2_struct_0(X1)|k15_lattice3(X1,X2)=k16_lattice3(X1,a_2_2_lattice3(X1,X2))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_97, plain, (k15_lattice3(X1,X2)=k15_lattice3(X1,a_2_5_lattice3(X1,X2))|v2_struct_0(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_98, negated_conjecture, (a_2_5_lattice3(esk9_0,esk2_2(k1_zfmisc_1(k1_xboole_0),X1))=k1_xboole_0|r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_35, c_0_93])).
cnf(c_0_99, plain, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(X2,a_2_6_lattice3(X1,esk1_2(k1_xboole_0,k1_xboole_0)))), inference(spm,[status(thm)],[c_0_94, c_0_79])).
fof(c_0_100, plain, ![X58, X59, X60, X62]:((m1_subset_1(esk10_3(X58,X59,X60),u1_struct_0(X58))|r3_lattices(X58,k15_lattice3(X58,X59),k15_lattice3(X58,X60))|(v2_struct_0(X58)|(~v10_lattices(X58)|(~v4_lattice3(X58)|~l3_lattices(X58)))))&((r2_hidden(esk10_3(X58,X59,X60),X59)|r3_lattices(X58,k15_lattice3(X58,X59),k15_lattice3(X58,X60))|(v2_struct_0(X58)|(~v10_lattices(X58)|(~v4_lattice3(X58)|~l3_lattices(X58)))))&(~m1_subset_1(X62,u1_struct_0(X58))|(~r3_lattices(X58,esk10_3(X58,X59,X60),X62)|~r2_hidden(X62,X60))|r3_lattices(X58,k15_lattice3(X58,X59),k15_lattice3(X58,X60))|(v2_struct_0(X58)|(~v10_lattices(X58)|(~v4_lattice3(X58)|~l3_lattices(X58))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t48_lattice3])])])])])])).
cnf(c_0_101, plain, (r3_lattice3(X1,k15_lattice3(X1,X2),a_2_2_lattice3(X1,X2))|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_102, negated_conjecture, (k15_lattice3(esk9_0,esk2_2(k1_zfmisc_1(k1_xboole_0),X1))=k15_lattice3(esk9_0,k1_xboole_0)|r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_81]), c_0_82]), c_0_83])]), c_0_84])).
cnf(c_0_103, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|r1_tarski(a_2_5_lattice3(esk9_0,esk1_2(k1_xboole_0,k1_xboole_0)),X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_88]), c_0_81]), c_0_82]), c_0_83])]), c_0_84])).
fof(c_0_104, plain, ![X63, X64, X65, X66]:(((r4_lattice3(X63,X65,X64)|X65!=k15_lattice3(X63,X64)|~m1_subset_1(X65,u1_struct_0(X63))|(v2_struct_0(X63)|(~v10_lattices(X63)|(~v4_lattice3(X63)|~l3_lattices(X63))))|(v2_struct_0(X63)|~l3_lattices(X63)))&(~m1_subset_1(X66,u1_struct_0(X63))|(~r4_lattice3(X63,X66,X64)|r1_lattices(X63,X65,X66))|X65!=k15_lattice3(X63,X64)|~m1_subset_1(X65,u1_struct_0(X63))|(v2_struct_0(X63)|(~v10_lattices(X63)|(~v4_lattice3(X63)|~l3_lattices(X63))))|(v2_struct_0(X63)|~l3_lattices(X63))))&((m1_subset_1(esk11_3(X63,X64,X65),u1_struct_0(X63))|~r4_lattice3(X63,X65,X64)|X65=k15_lattice3(X63,X64)|~m1_subset_1(X65,u1_struct_0(X63))|(v2_struct_0(X63)|(~v10_lattices(X63)|(~v4_lattice3(X63)|~l3_lattices(X63))))|(v2_struct_0(X63)|~l3_lattices(X63)))&((r4_lattice3(X63,esk11_3(X63,X64,X65),X64)|~r4_lattice3(X63,X65,X64)|X65=k15_lattice3(X63,X64)|~m1_subset_1(X65,u1_struct_0(X63))|(v2_struct_0(X63)|(~v10_lattices(X63)|(~v4_lattice3(X63)|~l3_lattices(X63))))|(v2_struct_0(X63)|~l3_lattices(X63)))&(~r1_lattices(X63,X65,esk11_3(X63,X64,X65))|~r4_lattice3(X63,X65,X64)|X65=k15_lattice3(X63,X64)|~m1_subset_1(X65,u1_struct_0(X63))|(v2_struct_0(X63)|(~v10_lattices(X63)|(~v4_lattice3(X63)|~l3_lattices(X63))))|(v2_struct_0(X63)|~l3_lattices(X63)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d21_lattice3])])])])])])).
cnf(c_0_105, plain, (r2_hidden(esk10_3(X1,X2,X3),X2)|r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3))|v2_struct_0(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_106, negated_conjecture, (r3_lattice3(esk9_0,k15_lattice3(esk9_0,k1_xboole_0),a_2_2_lattice3(esk9_0,esk2_2(k1_zfmisc_1(k1_xboole_0),X1)))|r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_81]), c_0_82]), c_0_83])]), c_0_84])).
cnf(c_0_107, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_33, c_0_32])).
fof(c_0_108, plain, ![X68, X69]:(v2_struct_0(X68)|~l3_lattices(X68)|m1_subset_1(k15_lattice3(X68,X69),u1_struct_0(X68))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k15_lattice3])])])).
cnf(c_0_109, negated_conjecture, (a_2_5_lattice3(esk9_0,esk1_2(k1_xboole_0,k1_xboole_0))=k1_xboole_0|k1_zfmisc_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_35, c_0_103])).
cnf(c_0_110, plain, (r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X1)|X2!=k15_lattice3(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_111, negated_conjecture, (r3_lattices(esk9_0,k15_lattice3(esk9_0,X1),k15_lattice3(esk9_0,X2))|r2_hidden(esk10_3(esk9_0,X1,X2),X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_81]), c_0_82]), c_0_83])]), c_0_84])).
cnf(c_0_112, negated_conjecture, (r3_lattice3(esk9_0,k15_lattice3(esk9_0,k1_xboole_0),a_2_2_lattice3(esk9_0,k1_xboole_0))|r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_106, c_0_39])).
cnf(c_0_113, plain, (~r2_hidden(X1,X1)), inference(spm,[status(thm)],[c_0_51, c_0_107])).
cnf(c_0_114, plain, (v2_struct_0(X1)|m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
cnf(c_0_115, negated_conjecture, (k15_lattice3(esk9_0,esk1_2(k1_xboole_0,k1_xboole_0))=k15_lattice3(esk9_0,k1_xboole_0)|k1_zfmisc_1(k1_xboole_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_109]), c_0_81]), c_0_82]), c_0_83])]), c_0_84])).
cnf(c_0_116, plain, (v2_struct_0(X1)|r4_lattice3(X1,X2,X3)|X2!=k15_lattice3(X1,X3)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_110])).
cnf(c_0_117, plain, (v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(X2,a_2_6_lattice3(X1,k1_xboole_0))), inference(spm,[status(thm)],[c_0_56, c_0_79])).
cnf(c_0_118, negated_conjecture, (r3_lattices(esk9_0,k15_lattice3(esk9_0,esk2_2(k1_zfmisc_1(k1_xboole_0),X1)),k15_lattice3(esk9_0,X2))|r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_66, c_0_111])).
fof(c_0_119, plain, ![X70, X71, X72]:(v2_struct_0(X70)|(~v10_lattices(X70)|(~v4_lattice3(X70)|~l3_lattices(X70)))|(~m1_subset_1(X71,u1_struct_0(X70))|(~r2_hidden(X71,X72)|~r3_lattice3(X70,X71,X72)|k16_lattice3(X70,X72)=X71))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t42_lattice3])])])])).
cnf(c_0_120, negated_conjecture, (r3_lattice3(esk9_0,k15_lattice3(esk9_0,k1_xboole_0),a_2_2_lattice3(esk9_0,k1_xboole_0))|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_51, c_0_112])).
cnf(c_0_121, plain, (~r1_tarski(k1_zfmisc_1(X1),X1)), inference(spm,[status(thm)],[c_0_113, c_0_34])).
cnf(c_0_122, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|m1_subset_1(k15_lattice3(esk9_0,k1_xboole_0),u1_struct_0(esk9_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_81])]), c_0_84])).
fof(c_0_123, plain, ![X73, X74, X75, X77, X78]:(((m1_subset_1(esk12_3(X73,X74,X75),u1_struct_0(X74))|~r2_hidden(X73,a_2_2_lattice3(X74,X75))|(v2_struct_0(X74)|(~v10_lattices(X74)|(~v4_lattice3(X74)|~l3_lattices(X74)))))&((X73=esk12_3(X73,X74,X75)|~r2_hidden(X73,a_2_2_lattice3(X74,X75))|(v2_struct_0(X74)|(~v10_lattices(X74)|(~v4_lattice3(X74)|~l3_lattices(X74)))))&(r4_lattice3(X74,esk12_3(X73,X74,X75),X75)|~r2_hidden(X73,a_2_2_lattice3(X74,X75))|(v2_struct_0(X74)|(~v10_lattices(X74)|(~v4_lattice3(X74)|~l3_lattices(X74)))))))&(~m1_subset_1(X78,u1_struct_0(X74))|(X73!=X78|~r4_lattice3(X74,X78,X77))|r2_hidden(X73,a_2_2_lattice3(X74,X77))|(v2_struct_0(X74)|(~v10_lattices(X74)|(~v4_lattice3(X74)|~l3_lattices(X74)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_2_lattice3])])])])])])])).
cnf(c_0_124, plain, (r4_lattice3(X1,k15_lattice3(X1,X2),X2)|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_116]), c_0_114])).
cnf(c_0_125, negated_conjecture, (r1_tarski(a_2_5_lattice3(esk9_0,k1_xboole_0),X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_88]), c_0_81]), c_0_82]), c_0_83])]), c_0_84])).
cnf(c_0_126, negated_conjecture, (r3_lattices(esk9_0,k15_lattice3(esk9_0,k1_xboole_0),k15_lattice3(esk9_0,X1))|r1_tarski(k1_zfmisc_1(k1_xboole_0),X2)), inference(spm,[status(thm)],[c_0_118, c_0_39])).
cnf(c_0_127, plain, (v2_struct_0(X1)|k16_lattice3(X1,X3)=X2|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X2,X3)|~r3_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_128, negated_conjecture, (r3_lattice3(esk9_0,k15_lattice3(esk9_0,k1_xboole_0),a_2_2_lattice3(esk9_0,k1_xboole_0))), inference(spm,[status(thm)],[c_0_120, c_0_86])).
cnf(c_0_129, negated_conjecture, (m1_subset_1(k15_lattice3(esk9_0,k1_xboole_0),u1_struct_0(esk9_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_40])])).
cnf(c_0_130, plain, (r2_hidden(X3,a_2_2_lattice3(X2,X4))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~r4_lattice3(X2,X1,X4)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_123])).
cnf(c_0_131, plain, (r4_lattice3(X1,k15_lattice3(X1,X2),a_2_5_lattice3(X1,X2))|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_124, c_0_97])).
cnf(c_0_132, negated_conjecture, (a_2_5_lattice3(esk9_0,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_35, c_0_125])).
cnf(c_0_133, negated_conjecture, (r3_lattices(esk9_0,k15_lattice3(esk9_0,k1_xboole_0),k15_lattice3(esk9_0,X1))|~r2_hidden(X2,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_51, c_0_126])).
fof(c_0_134, plain, ![X79, X80]:((k15_lattice3(X79,k6_domain_1(u1_struct_0(X79),X80))=X80|~m1_subset_1(X80,u1_struct_0(X79))|(v2_struct_0(X79)|(~v10_lattices(X79)|(~v4_lattice3(X79)|~l3_lattices(X79)))))&(k16_lattice3(X79,k6_domain_1(u1_struct_0(X79),X80))=X80|~m1_subset_1(X80,u1_struct_0(X79))|(v2_struct_0(X79)|(~v10_lattices(X79)|(~v4_lattice3(X79)|~l3_lattices(X79)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t43_lattice3])])])])])).
cnf(c_0_135, plain, (k16_lattice3(X1,X2)=k16_lattice3(X1,a_2_6_lattice3(X1,X2))|v2_struct_0(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_136, negated_conjecture, (k16_lattice3(esk9_0,a_2_2_lattice3(esk9_0,k1_xboole_0))=k15_lattice3(esk9_0,k1_xboole_0)|~r2_hidden(k15_lattice3(esk9_0,k1_xboole_0),a_2_2_lattice3(esk9_0,k1_xboole_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128]), c_0_129]), c_0_81]), c_0_82]), c_0_83])]), c_0_84])).
cnf(c_0_137, plain, (v2_struct_0(X1)|r2_hidden(X2,a_2_2_lattice3(X1,X3))|~r4_lattice3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(er,[status(thm)],[c_0_130])).
cnf(c_0_138, negated_conjecture, (r4_lattice3(esk9_0,k15_lattice3(esk9_0,k1_xboole_0),k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_132]), c_0_81]), c_0_82]), c_0_83])]), c_0_84])).
cnf(c_0_139, negated_conjecture, (r3_lattices(esk9_0,k15_lattice3(esk9_0,k1_xboole_0),k15_lattice3(esk9_0,X1))), inference(spm,[status(thm)],[c_0_133, c_0_86])).
cnf(c_0_140, plain, (k15_lattice3(X1,k6_domain_1(u1_struct_0(X1),X2))=X2|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_141, plain, (r3_lattice3(X1,k16_lattice3(X1,X2),a_2_6_lattice3(X1,X2))|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_95, c_0_135])).
cnf(c_0_142, negated_conjecture, (k16_lattice3(esk9_0,a_2_2_lattice3(esk9_0,k1_xboole_0))=k15_lattice3(esk9_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_137]), c_0_138]), c_0_129]), c_0_81]), c_0_82]), c_0_83])]), c_0_84])).
cnf(c_0_143, negated_conjecture, (r3_lattices(esk9_0,k15_lattice3(esk9_0,k1_xboole_0),X1)|~m1_subset_1(X1,u1_struct_0(esk9_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140]), c_0_81]), c_0_82]), c_0_83])]), c_0_84])).
cnf(c_0_144, plain, (r1_lattices(X1,X2,X4)|v2_struct_0(X1)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_145, negated_conjecture, (r3_lattice3(esk9_0,k15_lattice3(esk9_0,k1_xboole_0),a_2_6_lattice3(esk9_0,a_2_2_lattice3(esk9_0,k1_xboole_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_142]), c_0_81]), c_0_82]), c_0_83])]), c_0_84])).
cnf(c_0_146, negated_conjecture, (r2_hidden(X1,a_2_6_lattice3(esk9_0,X2))|~m1_subset_1(k15_lattice3(esk9_0,k1_xboole_0),u1_struct_0(esk9_0))|~m1_subset_1(X1,u1_struct_0(esk9_0))|~r2_hidden(k15_lattice3(esk9_0,k1_xboole_0),X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_143]), c_0_81]), c_0_82]), c_0_83])]), c_0_84])).
cnf(c_0_147, negated_conjecture, (r1_lattices(esk9_0,k15_lattice3(esk9_0,k1_xboole_0),X1)|~m1_subset_1(X1,u1_struct_0(esk9_0))|~r2_hidden(X1,a_2_6_lattice3(esk9_0,a_2_2_lattice3(esk9_0,k1_xboole_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_145]), c_0_129]), c_0_81])]), c_0_84])).
cnf(c_0_148, negated_conjecture, (r2_hidden(X1,a_2_6_lattice3(esk9_0,X2))|~m1_subset_1(X1,u1_struct_0(esk9_0))|~r2_hidden(k15_lattice3(esk9_0,k1_xboole_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_146, c_0_129])])).
cnf(c_0_149, negated_conjecture, (r1_lattices(esk9_0,k15_lattice3(esk9_0,k1_xboole_0),X1)|~m1_subset_1(X1,u1_struct_0(esk9_0))|~r2_hidden(k15_lattice3(esk9_0,k1_xboole_0),a_2_2_lattice3(esk9_0,k1_xboole_0))), inference(spm,[status(thm)],[c_0_147, c_0_148])).
cnf(c_0_150, plain, (r3_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r1_lattices(X1,X2,esk6_3(X1,X2,X3))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_151, negated_conjecture, (r1_lattices(esk9_0,k15_lattice3(esk9_0,k1_xboole_0),X1)|~m1_subset_1(X1,u1_struct_0(esk9_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_137]), c_0_138]), c_0_129]), c_0_81]), c_0_82]), c_0_83])]), c_0_84])).
cnf(c_0_152, negated_conjecture, (r3_lattice3(esk9_0,k15_lattice3(esk9_0,k1_xboole_0),X1)|~m1_subset_1(esk6_3(esk9_0,k15_lattice3(esk9_0,k1_xboole_0),X1),u1_struct_0(esk9_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_151]), c_0_129]), c_0_81])]), c_0_84])).
cnf(c_0_153, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|r3_lattice3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_154, negated_conjecture, (r3_lattice3(esk9_0,k15_lattice3(esk9_0,k1_xboole_0),X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_153]), c_0_129]), c_0_81])]), c_0_84])).
cnf(c_0_155, negated_conjecture, (r2_hidden(k16_lattice3(esk9_0,X1),a_2_6_lattice3(esk9_0,X2))|~r2_hidden(k15_lattice3(esk9_0,k1_xboole_0),X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_154]), c_0_129]), c_0_81]), c_0_82]), c_0_83])]), c_0_84])).
cnf(c_0_156, plain, (X1=esk3_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_6_lattice3(X2,X3))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_157, negated_conjecture, (r2_hidden(k15_lattice3(esk9_0,X1),a_2_6_lattice3(esk9_0,X2))|~r2_hidden(k15_lattice3(esk9_0,k1_xboole_0),X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_96]), c_0_81]), c_0_82]), c_0_83])]), c_0_84])).
cnf(c_0_158, plain, (m1_subset_1(esk3_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_6_lattice3(X2,X3))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_159, negated_conjecture, (esk3_3(k15_lattice3(esk9_0,X1),esk9_0,X2)=k15_lattice3(esk9_0,X1)|~r2_hidden(k15_lattice3(esk9_0,k1_xboole_0),X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_157]), c_0_81]), c_0_82]), c_0_83])]), c_0_84])).
fof(c_0_160, plain, ![X81, X82, X83]:(((k2_lattices(X81,X82,X83)=X82|~m1_subset_1(X83,u1_struct_0(X81))|X82!=k5_lattices(X81)|~m1_subset_1(X82,u1_struct_0(X81))|~v13_lattices(X81)|(v2_struct_0(X81)|~l1_lattices(X81)))&(k2_lattices(X81,X83,X82)=X82|~m1_subset_1(X83,u1_struct_0(X81))|X82!=k5_lattices(X81)|~m1_subset_1(X82,u1_struct_0(X81))|~v13_lattices(X81)|(v2_struct_0(X81)|~l1_lattices(X81))))&((m1_subset_1(esk13_2(X81,X82),u1_struct_0(X81))|X82=k5_lattices(X81)|~m1_subset_1(X82,u1_struct_0(X81))|~v13_lattices(X81)|(v2_struct_0(X81)|~l1_lattices(X81)))&(k2_lattices(X81,X82,esk13_2(X81,X82))!=X82|k2_lattices(X81,esk13_2(X81,X82),X82)!=X82|X82=k5_lattices(X81)|~m1_subset_1(X82,u1_struct_0(X81))|~v13_lattices(X81)|(v2_struct_0(X81)|~l1_lattices(X81))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattices])])])])])])).
fof(c_0_161, plain, ![X85]:(v2_struct_0(X85)|~l1_lattices(X85)|m1_subset_1(k5_lattices(X85),u1_struct_0(X85))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_lattices])])])).
fof(c_0_162, plain, ![X86, X87, X88]:((~r1_lattices(X86,X87,X88)|k2_lattices(X86,X87,X88)=X87|~m1_subset_1(X88,u1_struct_0(X86))|~m1_subset_1(X87,u1_struct_0(X86))|(v2_struct_0(X86)|(~v8_lattices(X86)|(~v9_lattices(X86)|~l3_lattices(X86)))))&(k2_lattices(X86,X87,X88)!=X87|r1_lattices(X86,X87,X88)|~m1_subset_1(X88,u1_struct_0(X86))|~m1_subset_1(X87,u1_struct_0(X86))|(v2_struct_0(X86)|(~v8_lattices(X86)|(~v9_lattices(X86)|~l3_lattices(X86)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t21_lattices])])])])])).
cnf(c_0_163, negated_conjecture, (m1_subset_1(k15_lattice3(esk9_0,X1),u1_struct_0(esk9_0))|~r2_hidden(k15_lattice3(esk9_0,k1_xboole_0),X2)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159]), c_0_81]), c_0_82]), c_0_83])]), c_0_84]), c_0_157])).
cnf(c_0_164, plain, (k2_lattices(X1,X2,X3)=X3|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|X3!=k5_lattices(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v13_lattices(X1)|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_160])).
cnf(c_0_165, plain, (v2_struct_0(X1)|m1_subset_1(k5_lattices(X1),u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_161])).
cnf(c_0_166, plain, (k2_lattices(X1,X2,X3)=X2|v2_struct_0(X1)|~r1_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_162])).
cnf(c_0_167, negated_conjecture, (m1_subset_1(k15_lattice3(esk9_0,X1),u1_struct_0(esk9_0))|~r1_tarski(k15_lattice3(esk9_0,k1_xboole_0),X2)), inference(spm,[status(thm)],[c_0_163, c_0_34])).
cnf(c_0_168, negated_conjecture, (v2_struct_0(esk9_0)|~v10_lattices(esk9_0)|~v13_lattices(esk9_0)|~l3_lattices(esk9_0)|k5_lattices(esk9_0)!=k15_lattice3(esk9_0,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_169, plain, (k2_lattices(X1,X2,k5_lattices(X1))=k5_lattices(X1)|v2_struct_0(X1)|~l1_lattices(X1)|~v13_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_164]), c_0_165])).
cnf(c_0_170, negated_conjecture, (k2_lattices(esk9_0,k15_lattice3(esk9_0,k1_xboole_0),X1)=k15_lattice3(esk9_0,k1_xboole_0)|~v9_lattices(esk9_0)|~v8_lattices(esk9_0)|~m1_subset_1(X1,u1_struct_0(esk9_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_151]), c_0_129]), c_0_81])]), c_0_84])).
cnf(c_0_171, negated_conjecture, (m1_subset_1(k15_lattice3(esk9_0,X1),u1_struct_0(esk9_0))), inference(spm,[status(thm)],[c_0_167, c_0_107])).
cnf(c_0_172, negated_conjecture, (k5_lattices(esk9_0)!=k15_lattice3(esk9_0,k1_xboole_0)|~v13_lattices(esk9_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_168, c_0_83]), c_0_81])]), c_0_84])).
fof(c_0_173, plain, ![X89, X90, X91]:((~v6_lattices(X89)|(~m1_subset_1(X90,u1_struct_0(X89))|(~m1_subset_1(X91,u1_struct_0(X89))|k2_lattices(X89,X90,X91)=k2_lattices(X89,X91,X90)))|(v2_struct_0(X89)|~l1_lattices(X89)))&((m1_subset_1(esk14_1(X89),u1_struct_0(X89))|v6_lattices(X89)|(v2_struct_0(X89)|~l1_lattices(X89)))&((m1_subset_1(esk15_1(X89),u1_struct_0(X89))|v6_lattices(X89)|(v2_struct_0(X89)|~l1_lattices(X89)))&(k2_lattices(X89,esk14_1(X89),esk15_1(X89))!=k2_lattices(X89,esk15_1(X89),esk14_1(X89))|v6_lattices(X89)|(v2_struct_0(X89)|~l1_lattices(X89)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_lattices])])])])])])).
fof(c_0_174, plain, ![X94]:((~v2_struct_0(X94)|(v2_struct_0(X94)|~v10_lattices(X94))|~l3_lattices(X94))&((v4_lattices(X94)|(v2_struct_0(X94)|~v10_lattices(X94))|~l3_lattices(X94))&((v5_lattices(X94)|(v2_struct_0(X94)|~v10_lattices(X94))|~l3_lattices(X94))&((v6_lattices(X94)|(v2_struct_0(X94)|~v10_lattices(X94))|~l3_lattices(X94))&((v7_lattices(X94)|(v2_struct_0(X94)|~v10_lattices(X94))|~l3_lattices(X94))&((v8_lattices(X94)|(v2_struct_0(X94)|~v10_lattices(X94))|~l3_lattices(X94))&(v9_lattices(X94)|(v2_struct_0(X94)|~v10_lattices(X94))|~l3_lattices(X94)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
fof(c_0_175, plain, ![X95]:((l1_lattices(X95)|~l3_lattices(X95))&(l2_lattices(X95)|~l3_lattices(X95))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
cnf(c_0_176, negated_conjecture, (~v9_lattices(esk9_0)|~v8_lattices(esk9_0)|~l1_lattices(esk9_0)|~v13_lattices(esk9_0)|~m1_subset_1(k5_lattices(esk9_0),u1_struct_0(esk9_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_170]), c_0_171])]), c_0_84]), c_0_172])).
cnf(c_0_177, plain, (k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2)|v2_struct_0(X1)|~v6_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_173])).
cnf(c_0_178, plain, (v6_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_179, plain, (l1_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_175])).
cnf(c_0_180, negated_conjecture, (~v9_lattices(esk9_0)|~v8_lattices(esk9_0)|~l1_lattices(esk9_0)|~v13_lattices(esk9_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_165]), c_0_84])).
cnf(c_0_181, plain, (v9_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_174])).
fof(c_0_182, plain, ![X96, X98, X99]:(((m1_subset_1(esk16_1(X96),u1_struct_0(X96))|~v13_lattices(X96)|(v2_struct_0(X96)|~l1_lattices(X96)))&((k2_lattices(X96,esk16_1(X96),X98)=esk16_1(X96)|~m1_subset_1(X98,u1_struct_0(X96))|~v13_lattices(X96)|(v2_struct_0(X96)|~l1_lattices(X96)))&(k2_lattices(X96,X98,esk16_1(X96))=esk16_1(X96)|~m1_subset_1(X98,u1_struct_0(X96))|~v13_lattices(X96)|(v2_struct_0(X96)|~l1_lattices(X96)))))&((m1_subset_1(esk17_2(X96,X99),u1_struct_0(X96))|~m1_subset_1(X99,u1_struct_0(X96))|v13_lattices(X96)|(v2_struct_0(X96)|~l1_lattices(X96)))&(k2_lattices(X96,X99,esk17_2(X96,X99))!=X99|k2_lattices(X96,esk17_2(X96,X99),X99)!=X99|~m1_subset_1(X99,u1_struct_0(X96))|v13_lattices(X96)|(v2_struct_0(X96)|~l1_lattices(X96))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_lattices])])])])])])).
cnf(c_0_183, plain, (k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_178]), c_0_179])).
cnf(c_0_184, negated_conjecture, (~v8_lattices(esk9_0)|~l1_lattices(esk9_0)|~v13_lattices(esk9_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_181]), c_0_81]), c_0_83])]), c_0_84])).
cnf(c_0_185, plain, (v8_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_186, plain, (v13_lattices(X1)|v2_struct_0(X1)|k2_lattices(X1,X2,esk17_2(X1,X2))!=X2|k2_lattices(X1,esk17_2(X1,X2),X2)!=X2|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_182])).
cnf(c_0_187, negated_conjecture, (k2_lattices(esk9_0,X1,k15_lattice3(esk9_0,k1_xboole_0))=k15_lattice3(esk9_0,k1_xboole_0)|~v9_lattices(esk9_0)|~v8_lattices(esk9_0)|~m1_subset_1(X1,u1_struct_0(esk9_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_170]), c_0_171]), c_0_81]), c_0_83])]), c_0_84])).
cnf(c_0_188, negated_conjecture, (~l1_lattices(esk9_0)|~v13_lattices(esk9_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_185]), c_0_81]), c_0_83])]), c_0_84])).
cnf(c_0_189, negated_conjecture, (~v9_lattices(esk9_0)|~v8_lattices(esk9_0)|~l1_lattices(esk9_0)|~m1_subset_1(esk17_2(esk9_0,k15_lattice3(esk9_0,k1_xboole_0)),u1_struct_0(esk9_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_187]), c_0_171])]), c_0_84]), c_0_170]), c_0_188])).
cnf(c_0_190, plain, (m1_subset_1(esk17_2(X1,X2),u1_struct_0(X1))|v13_lattices(X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_182])).
cnf(c_0_191, negated_conjecture, (~v9_lattices(esk9_0)|~v8_lattices(esk9_0)|~l1_lattices(esk9_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_190]), c_0_171])]), c_0_84]), c_0_188])).
cnf(c_0_192, negated_conjecture, (~v8_lattices(esk9_0)|~l1_lattices(esk9_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_181]), c_0_81]), c_0_83])]), c_0_84])).
cnf(c_0_193, negated_conjecture, (~l1_lattices(esk9_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_185]), c_0_81]), c_0_83])]), c_0_84])).
cnf(c_0_194, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_179]), c_0_81])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 195
# Proof object clause steps            : 142
# Proof object formula steps           : 53
# Proof object conjectures             : 58
# Proof object clause conjectures      : 55
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 46
# Proof object initial formulas used   : 26
# Proof object generating inferences   : 85
# Proof object simplifying inferences  : 181
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 26
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 81
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 80
# Processed clauses                    : 5817
# ...of these trivial                  : 9
# ...subsumed                          : 4216
# ...remaining for further processing  : 1592
# Other redundant clauses eliminated   : 11
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 194
# Backward-rewritten                   : 100
# Generated clauses                    : 31184
# ...of the previous two non-trivial   : 29564
# Contextual simplify-reflections      : 117
# Paramodulations                      : 31157
# Factorizations                       : 16
# Equation resolutions                 : 11
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1207
#    Positive orientable unit clauses  : 28
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 1172
# Current number of unprocessed clauses: 23306
# ...number of literals in the above   : 174004
# Current number of archived formulas  : 0
# Current number of archived clauses   : 374
# Clause-clause subsumption calls (NU) : 957850
# Rec. Clause-clause subsumption calls : 143619
# Non-unit clause-clause subsumptions  : 3884
# Unit Clause-clause subsumption calls : 2007
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 128
# BW rewrite match successes           : 39
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 837985

# -------------------------------------------------
# User time                : 0.667 s
# System time              : 0.000 s
# Total time               : 0.667 s
# Maximum resident set size: 30136 pages
