# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/0b5d7f60e26f4890a77017ebf12a4729.E_input', d5_xboole_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/0b5d7f60e26f4890a77017ebf12a4729.E_input', d1_xboole_0)).
fof(t9_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r1_waybel_0(X1,X2,X3)<=>~(r2_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3)))))), file('data/MPTP/proofs/0b5d7f60e26f4890a77017ebf12a4729.E_input', t9_waybel_0)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/0b5d7f60e26f4890a77017ebf12a4729.E_input', redefinition_k6_subset_1)).
fof(t29_yellow_6, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>r1_waybel_0(X1,X2,u1_struct_0(X1)))), file('data/MPTP/proofs/0b5d7f60e26f4890a77017ebf12a4729.E_input', t29_yellow_6)).
fof(d12_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r2_waybel_0(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X4,X5)&r2_hidden(k2_waybel_0(X1,X2,X5),X3))))))), file('data/MPTP/proofs/0b5d7f60e26f4890a77017ebf12a4729.E_input', d12_waybel_0)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/0b5d7f60e26f4890a77017ebf12a4729.E_input', d2_subset_1)).
fof(c_0_7, plain, ![X6, X7, X8, X9, X10, X11, X12, X13]:((((r2_hidden(X9,X6)|~r2_hidden(X9,X8)|X8!=k4_xboole_0(X6,X7))&(~r2_hidden(X9,X7)|~r2_hidden(X9,X8)|X8!=k4_xboole_0(X6,X7)))&(~r2_hidden(X10,X6)|r2_hidden(X10,X7)|r2_hidden(X10,X8)|X8!=k4_xboole_0(X6,X7)))&((~r2_hidden(esk1_3(X11,X12,X13),X13)|(~r2_hidden(esk1_3(X11,X12,X13),X11)|r2_hidden(esk1_3(X11,X12,X13),X12))|X13=k4_xboole_0(X11,X12))&((r2_hidden(esk1_3(X11,X12,X13),X11)|r2_hidden(esk1_3(X11,X12,X13),X13)|X13=k4_xboole_0(X11,X12))&(~r2_hidden(esk1_3(X11,X12,X13),X12)|r2_hidden(esk1_3(X11,X12,X13),X13)|X13=k4_xboole_0(X11,X12))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_8, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_9, plain, ![X17, X18, X19]:((~v1_xboole_0(X17)|~r2_hidden(X18,X17))&(r2_hidden(esk2_1(X19),X19)|v1_xboole_0(X19))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_10, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (r2_hidden(esk1_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk1_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (r2_hidden(esk1_3(X1,X2,X3),X1)|r2_hidden(esk1_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_8])).
cnf(c_0_14, plain, (r2_hidden(esk2_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, plain, (X1=k4_xboole_0(X2,X2)|r2_hidden(esk1_3(X2,X2,X1),X1)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_18, plain, (v1_xboole_0(k4_xboole_0(X1,X2))|~r2_hidden(esk2_1(k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_19, plain, (v1_xboole_0(k4_xboole_0(X1,X2))|r2_hidden(esk2_1(k4_xboole_0(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_15, c_0_14])).
cnf(c_0_20, plain, (X1=k4_xboole_0(X2,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_21, plain, (v1_xboole_0(k4_xboole_0(X1,X1))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
fof(c_0_22, plain, ![X21, X22, X23]:((~r1_waybel_0(X21,X22,X23)|~r2_waybel_0(X21,X22,k6_subset_1(u1_struct_0(X21),X23))|(v2_struct_0(X22)|~l1_waybel_0(X22,X21))|(v2_struct_0(X21)|~l1_struct_0(X21)))&(r2_waybel_0(X21,X22,k6_subset_1(u1_struct_0(X21),X23))|r1_waybel_0(X21,X22,X23)|(v2_struct_0(X22)|~l1_waybel_0(X22,X21))|(v2_struct_0(X21)|~l1_struct_0(X21)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_waybel_0])])])])])).
fof(c_0_23, plain, ![X15, X16]:k6_subset_1(X15,X16)=k4_xboole_0(X15,X16), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
cnf(c_0_24, plain, (k4_xboole_0(X1,X1)=k4_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_25, plain, (r2_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3))|r1_waybel_0(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_27, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>r1_waybel_0(X1,X2,u1_struct_0(X1))))), inference(assume_negation,[status(cth)],[t29_yellow_6])).
cnf(c_0_28, plain, (~r2_hidden(X1,k4_xboole_0(X2,X2))|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_13, c_0_24])).
fof(c_0_29, plain, ![X26, X27, X28, X29, X31, X33]:(((m1_subset_1(esk5_4(X26,X27,X28,X29),u1_struct_0(X27))|~m1_subset_1(X29,u1_struct_0(X27))|~r2_waybel_0(X26,X27,X28)|(v2_struct_0(X27)|~l1_waybel_0(X27,X26))|(v2_struct_0(X26)|~l1_struct_0(X26)))&((r1_orders_2(X27,X29,esk5_4(X26,X27,X28,X29))|~m1_subset_1(X29,u1_struct_0(X27))|~r2_waybel_0(X26,X27,X28)|(v2_struct_0(X27)|~l1_waybel_0(X27,X26))|(v2_struct_0(X26)|~l1_struct_0(X26)))&(r2_hidden(k2_waybel_0(X26,X27,esk5_4(X26,X27,X28,X29)),X28)|~m1_subset_1(X29,u1_struct_0(X27))|~r2_waybel_0(X26,X27,X28)|(v2_struct_0(X27)|~l1_waybel_0(X27,X26))|(v2_struct_0(X26)|~l1_struct_0(X26)))))&((m1_subset_1(esk6_3(X26,X27,X31),u1_struct_0(X27))|r2_waybel_0(X26,X27,X31)|(v2_struct_0(X27)|~l1_waybel_0(X27,X26))|(v2_struct_0(X26)|~l1_struct_0(X26)))&(~m1_subset_1(X33,u1_struct_0(X27))|(~r1_orders_2(X27,esk6_3(X26,X27,X31),X33)|~r2_hidden(k2_waybel_0(X26,X27,X33),X31))|r2_waybel_0(X26,X27,X31)|(v2_struct_0(X27)|~l1_waybel_0(X27,X26))|(v2_struct_0(X26)|~l1_struct_0(X26))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d12_waybel_0])])])])])])])).
cnf(c_0_30, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r1_waybel_0(X1,X2,X3)|r2_waybel_0(X1,X2,k4_xboole_0(u1_struct_0(X1),X3))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)), inference(rw,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_31, negated_conjecture, ((~v2_struct_0(esk3_0)&l1_struct_0(esk3_0))&((~v2_struct_0(esk4_0)&(v4_orders_2(esk4_0)&(v7_waybel_0(esk4_0)&l1_waybel_0(esk4_0,esk3_0))))&~r1_waybel_0(esk3_0,esk4_0,u1_struct_0(esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_27])])])])).
cnf(c_0_32, plain, (~r2_hidden(X1,k4_xboole_0(X2,X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_24]), c_0_28])).
cnf(c_0_33, plain, (r2_hidden(k2_waybel_0(X1,X2,esk5_4(X1,X2,X3,X4)),X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~r2_waybel_0(X1,X2,X3)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34, plain, (r2_waybel_0(X1,X2,k4_xboole_0(X3,X3))|r1_waybel_0(X1,X2,u1_struct_0(X1))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_30, c_0_24])).
cnf(c_0_35, negated_conjecture, (l1_waybel_0(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_36, negated_conjecture, (l1_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_37, negated_conjecture, (~r1_waybel_0(esk3_0,esk4_0,u1_struct_0(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_39, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_40, plain, ![X34, X35]:(((~m1_subset_1(X35,X34)|r2_hidden(X35,X34)|v1_xboole_0(X34))&(~r2_hidden(X35,X34)|m1_subset_1(X35,X34)|v1_xboole_0(X34)))&((~m1_subset_1(X35,X34)|v1_xboole_0(X35)|~v1_xboole_0(X34))&(~v1_xboole_0(X35)|m1_subset_1(X35,X34)|~v1_xboole_0(X34)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
cnf(c_0_41, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_waybel_0(X1,X2,k4_xboole_0(X4,X4))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_42, negated_conjecture, (r2_waybel_0(esk3_0,esk4_0,k4_xboole_0(X1,X1))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])]), c_0_37]), c_0_38]), c_0_39])).
cnf(c_0_43, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_44, negated_conjecture, (~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_35]), c_0_36])]), c_0_38]), c_0_39])).
cnf(c_0_45, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_43, c_0_16])).
cnf(c_0_46, negated_conjecture, (~r2_hidden(X1,u1_struct_0(esk4_0))), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_47, plain, (m1_subset_1(X1,X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_48, negated_conjecture, (v1_xboole_0(u1_struct_0(esk4_0))), inference(spm,[status(thm)],[c_0_46, c_0_14])).
cnf(c_0_49, negated_conjecture, (~v1_xboole_0(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_44])).
cnf(c_0_50, plain, ($false), inference(sr,[status(thm)],[c_0_21, c_0_49]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 51
# Proof object clause steps            : 36
# Proof object formula steps           : 15
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 26
# Processed clauses                    : 662
# ...of these trivial                  : 6
# ...subsumed                          : 502
# ...remaining for further processing  : 154
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 3
# Generated clauses                    : 1896
# ...of the previous two non-trivial   : 1541
# Contextual simplify-reflections      : 4
# Paramodulations                      : 1878
# Factorizations                       : 8
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 112
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 10
#    Non-unit-clauses                  : 90
# Current number of unprocessed clauses: 920
# ...number of literals in the above   : 3220
# Current number of archived formulas  : 0
# Current number of archived clauses   : 40
# Clause-clause subsumption calls (NU) : 2366
# Rec. Clause-clause subsumption calls : 1577
# Non-unit clause-clause subsumptions  : 294
# Unit Clause-clause subsumption calls : 208
# Rewrite failures with RHS unbound    : 79
# BW rewrite match attempts            : 68
# BW rewrite match successes           : 12
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 21920

# -------------------------------------------------
# User time                : 0.022 s
# System time              : 0.003 s
# Total time               : 0.026 s
# Maximum resident set size: 28720 pages
