# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t190_funct_2, conjecture, ![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X2,X3)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))=>~((r2_hidden(X1,k2_relset_1(X2,X3,X4))&![X5]:(m1_subset_1(X5,X2)=>~(X1=k1_funct_1(X4,X5)))))), file('data/MPTP/proofs/0b8a6690278f45358b2321b7a29ba66a.E_input', t190_funct_2)).
fof(d12_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2, X3]:(X3=k9_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(X5,k1_relat_1(X1))&(r2_hidden(X5,X2)&X4=k1_funct_1(X1,X5)))))), file('data/MPTP/proofs/0b8a6690278f45358b2321b7a29ba66a.E_input', d12_funct_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/0b8a6690278f45358b2321b7a29ba66a.E_input', cc1_relset_1)).
fof(t38_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(k7_relset_1(X1,X2,X3,X1)=k2_relset_1(X1,X2,X3)&k8_relset_1(X1,X2,X3,X2)=k1_relset_1(X1,X2,X3))), file('data/MPTP/proofs/0b8a6690278f45358b2321b7a29ba66a.E_input', t38_relset_1)).
fof(redefinition_k7_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k7_relset_1(X1,X2,X3,X4)=k9_relat_1(X3,X4)), file('data/MPTP/proofs/0b8a6690278f45358b2321b7a29ba66a.E_input', redefinition_k7_relset_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/0b8a6690278f45358b2321b7a29ba66a.E_input', t1_subset)).
fof(c_0_6, negated_conjecture, ~(![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X2,X3)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))=>~((r2_hidden(X1,k2_relset_1(X2,X3,X4))&![X5]:(m1_subset_1(X5,X2)=>~(X1=k1_funct_1(X4,X5))))))), inference(assume_negation,[status(cth)],[t190_funct_2])).
fof(c_0_7, plain, ![X11, X12, X13, X14, X16, X17, X18, X19, X21]:((((r2_hidden(esk5_4(X11,X12,X13,X14),k1_relat_1(X11))|~r2_hidden(X14,X13)|X13!=k9_relat_1(X11,X12)|(~v1_relat_1(X11)|~v1_funct_1(X11)))&((r2_hidden(esk5_4(X11,X12,X13,X14),X12)|~r2_hidden(X14,X13)|X13!=k9_relat_1(X11,X12)|(~v1_relat_1(X11)|~v1_funct_1(X11)))&(X14=k1_funct_1(X11,esk5_4(X11,X12,X13,X14))|~r2_hidden(X14,X13)|X13!=k9_relat_1(X11,X12)|(~v1_relat_1(X11)|~v1_funct_1(X11)))))&(~r2_hidden(X17,k1_relat_1(X11))|(~r2_hidden(X17,X12)|X16!=k1_funct_1(X11,X17))|r2_hidden(X16,X13)|X13!=k9_relat_1(X11,X12)|(~v1_relat_1(X11)|~v1_funct_1(X11))))&((~r2_hidden(esk6_3(X11,X18,X19),X19)|(~r2_hidden(X21,k1_relat_1(X11))|(~r2_hidden(X21,X18)|esk6_3(X11,X18,X19)!=k1_funct_1(X11,X21)))|X19=k9_relat_1(X11,X18)|(~v1_relat_1(X11)|~v1_funct_1(X11)))&((r2_hidden(esk7_3(X11,X18,X19),k1_relat_1(X11))|r2_hidden(esk6_3(X11,X18,X19),X19)|X19=k9_relat_1(X11,X18)|(~v1_relat_1(X11)|~v1_funct_1(X11)))&((r2_hidden(esk7_3(X11,X18,X19),X18)|r2_hidden(esk6_3(X11,X18,X19),X19)|X19=k9_relat_1(X11,X18)|(~v1_relat_1(X11)|~v1_funct_1(X11)))&(esk6_3(X11,X18,X19)=k1_funct_1(X11,esk7_3(X11,X18,X19))|r2_hidden(esk6_3(X11,X18,X19),X19)|X19=k9_relat_1(X11,X18)|(~v1_relat_1(X11)|~v1_funct_1(X11))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_funct_1])])])])])])).
fof(c_0_8, plain, ![X23, X24, X25]:(~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(X23,X24)))|v1_relat_1(X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_9, negated_conjecture, ![X10]:((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,esk2_0,esk3_0)&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))))&(r2_hidden(esk1_0,k2_relset_1(esk2_0,esk3_0,esk4_0))&(~m1_subset_1(X10,esk2_0)|esk1_0!=k1_funct_1(esk4_0,X10)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_6])])])])])).
fof(c_0_10, plain, ![X26, X27, X28]:((k7_relset_1(X26,X27,X28,X26)=k2_relset_1(X26,X27,X28)|~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(X26,X27))))&(k8_relset_1(X26,X27,X28,X27)=k1_relset_1(X26,X27,X28)|~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(X26,X27))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t38_relset_1])])])).
fof(c_0_11, plain, ![X29, X30, X31, X32]:(~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(X29,X30)))|k7_relset_1(X29,X30,X31,X32)=k9_relat_1(X31,X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_relset_1])])).
cnf(c_0_12, plain, (X1=k1_funct_1(X2,esk5_4(X2,X3,X4,X1))|~r2_hidden(X1,X4)|X4!=k9_relat_1(X2,X3)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (k7_relset_1(X1,X2,X3,X1)=k2_relset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (k7_relset_1(X2,X3,X1,X4)=k9_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, negated_conjecture, (~m1_subset_1(X1,esk2_0)|esk1_0!=k1_funct_1(esk4_0,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_18, plain, (k1_funct_1(X1,esk5_4(X1,X2,k9_relat_1(X1,X2),X3))=X3|~v1_relat_1(X1)|~r2_hidden(X3,k9_relat_1(X1,X2))|~v1_funct_1(X1)), inference(er,[status(thm)],[c_0_12])).
cnf(c_0_19, negated_conjecture, (v1_relat_1(esk4_0)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_20, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_21, plain, ![X33, X34]:(~r2_hidden(X33,X34)|m1_subset_1(X33,X34)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_22, negated_conjecture, (k7_relset_1(esk2_0,esk3_0,esk4_0,esk2_0)=k2_relset_1(esk2_0,esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_15, c_0_14])).
cnf(c_0_23, negated_conjecture, (k7_relset_1(esk2_0,esk3_0,esk4_0,X1)=k9_relat_1(esk4_0,X1)), inference(spm,[status(thm)],[c_0_16, c_0_14])).
cnf(c_0_24, negated_conjecture, (~r2_hidden(esk1_0,k9_relat_1(esk4_0,X1))|~m1_subset_1(esk5_4(esk4_0,X1,k9_relat_1(esk4_0,X1),esk1_0),esk2_0)), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19]), c_0_20])])])).
cnf(c_0_25, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_26, plain, (r2_hidden(esk5_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k9_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_27, negated_conjecture, (r2_hidden(esk1_0,k2_relset_1(esk2_0,esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_28, negated_conjecture, (k2_relset_1(esk2_0,esk3_0,esk4_0)=k9_relat_1(esk4_0,esk2_0)), inference(rw,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_29, negated_conjecture, (~r2_hidden(esk5_4(esk4_0,X1,k9_relat_1(esk4_0,X1),esk1_0),esk2_0)|~r2_hidden(esk1_0,k9_relat_1(esk4_0,X1))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_30, plain, (r2_hidden(esk5_4(X1,X2,k9_relat_1(X1,X2),X3),X2)|~v1_relat_1(X1)|~r2_hidden(X3,k9_relat_1(X1,X2))|~v1_funct_1(X1)), inference(er,[status(thm)],[c_0_26])).
cnf(c_0_31, negated_conjecture, (r2_hidden(esk1_0,k9_relat_1(esk4_0,esk2_0))), inference(rw,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_32, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_19]), c_0_20])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 33
# Proof object clause steps            : 20
# Proof object formula steps           : 13
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 18
# Processed clauses                    : 49
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 49
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 21
# ...of the previous two non-trivial   : 21
# Contextual simplify-reflections      : 0
# Paramodulations                      : 16
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 25
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 17
# Current number of unprocessed clauses: 8
# ...number of literals in the above   : 25
# Current number of archived formulas  : 0
# Current number of archived clauses   : 20
# Clause-clause subsumption calls (NU) : 100
# Rec. Clause-clause subsumption calls : 21
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1895

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.003 s
# Total time               : 0.009 s
# Maximum resident set size: 26280 pages
