# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t106_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=k6_tmap_1(X1,X2)))), file('data/MPTP/proofs/0b958edeff514fc6aa4fc168244a9b60.E_input', t106_tmap_1)).
fof(t1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/0b958edeff514fc6aa4fc168244a9b60.E_input', t1_tsep_1)).
fof(t116_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>((v1_tsep_1(X2,X1)&m1_pre_topc(X2,X1))<=>g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=k8_tmap_1(X1,X2)))), file('data/MPTP/proofs/0b958edeff514fc6aa4fc168244a9b60.E_input', t116_tmap_1)).
fof(d11_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:((v1_pre_topc(X3)&(v2_pre_topc(X3)&l1_pre_topc(X3)))=>(X3=k8_tmap_1(X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(X4=u1_struct_0(X2)=>X3=k6_tmap_1(X1,X4))))))), file('data/MPTP/proofs/0b958edeff514fc6aa4fc168244a9b60.E_input', d11_tmap_1)).
fof(dt_k8_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_pre_topc(X2,X1))=>(v1_pre_topc(k8_tmap_1(X1,X2))&(v2_pre_topc(k8_tmap_1(X1,X2))&l1_pre_topc(k8_tmap_1(X1,X2))))), file('data/MPTP/proofs/0b958edeff514fc6aa4fc168244a9b60.E_input', dt_k8_tmap_1)).
fof(d1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>(v1_tsep_1(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>v3_pre_topc(X3,X1)))))), file('data/MPTP/proofs/0b958edeff514fc6aa4fc168244a9b60.E_input', d1_tsep_1)).
fof(c_0_6, plain, ![X5, X6]:((~v3_pre_topc(X6,X5)|g1_pre_topc(u1_struct_0(X5),u1_pre_topc(X5))=k6_tmap_1(X5,X6)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&(g1_pre_topc(u1_struct_0(X5),u1_pre_topc(X5))!=k6_tmap_1(X5,X6)|v3_pre_topc(X6,X5)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t106_tmap_1])])])])])).
fof(c_0_7, plain, ![X7, X8]:(~l1_pre_topc(X7)|(~m1_pre_topc(X8,X7)|m1_subset_1(u1_struct_0(X8),k1_zfmisc_1(u1_struct_0(X7))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])).
fof(c_0_8, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>((v1_tsep_1(X2,X1)&m1_pre_topc(X2,X1))<=>g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=k8_tmap_1(X1,X2))))), inference(assume_negation,[status(cth)],[t116_tmap_1])).
cnf(c_0_9, plain, (v3_pre_topc(X2,X1)|v2_struct_0(X1)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k6_tmap_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_11, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((~v1_tsep_1(esk2_0,esk1_0)|~m1_pre_topc(esk2_0,esk1_0)|g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))!=k8_tmap_1(esk1_0,esk2_0))&((v1_tsep_1(esk2_0,esk1_0)|g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))=k8_tmap_1(esk1_0,esk2_0))&(m1_pre_topc(esk2_0,esk1_0)|g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))=k8_tmap_1(esk1_0,esk2_0)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])).
fof(c_0_12, plain, ![X11, X12, X13, X14]:((X13!=k8_tmap_1(X11,X12)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X11)))|(X14!=u1_struct_0(X12)|X13=k6_tmap_1(X11,X14)))|(~v1_pre_topc(X13)|(~v2_pre_topc(X13)|~l1_pre_topc(X13)))|~m1_pre_topc(X12,X11)|(v2_struct_0(X11)|(~v2_pre_topc(X11)|~l1_pre_topc(X11))))&((m1_subset_1(esk3_3(X11,X12,X13),k1_zfmisc_1(u1_struct_0(X11)))|X13=k8_tmap_1(X11,X12)|(~v1_pre_topc(X13)|(~v2_pre_topc(X13)|~l1_pre_topc(X13)))|~m1_pre_topc(X12,X11)|(v2_struct_0(X11)|(~v2_pre_topc(X11)|~l1_pre_topc(X11))))&((esk3_3(X11,X12,X13)=u1_struct_0(X12)|X13=k8_tmap_1(X11,X12)|(~v1_pre_topc(X13)|(~v2_pre_topc(X13)|~l1_pre_topc(X13)))|~m1_pre_topc(X12,X11)|(v2_struct_0(X11)|(~v2_pre_topc(X11)|~l1_pre_topc(X11))))&(X13!=k6_tmap_1(X11,esk3_3(X11,X12,X13))|X13=k8_tmap_1(X11,X12)|(~v1_pre_topc(X13)|(~v2_pre_topc(X13)|~l1_pre_topc(X13)))|~m1_pre_topc(X12,X11)|(v2_struct_0(X11)|(~v2_pre_topc(X11)|~l1_pre_topc(X11))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d11_tmap_1])])])])])])).
fof(c_0_13, plain, ![X16, X17]:((v1_pre_topc(k8_tmap_1(X16,X17))|(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))|~m1_pre_topc(X17,X16)))&((v2_pre_topc(k8_tmap_1(X16,X17))|(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))|~m1_pre_topc(X17,X16)))&(l1_pre_topc(k8_tmap_1(X16,X17))|(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))|~m1_pre_topc(X17,X16))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k8_tmap_1])])])])).
fof(c_0_14, plain, ![X18, X19, X20]:((~v1_tsep_1(X19,X18)|(~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X18)))|(X20!=u1_struct_0(X19)|v3_pre_topc(X20,X18)))|~m1_pre_topc(X19,X18)|~l1_pre_topc(X18))&((m1_subset_1(esk4_2(X18,X19),k1_zfmisc_1(u1_struct_0(X18)))|v1_tsep_1(X19,X18)|~m1_pre_topc(X19,X18)|~l1_pre_topc(X18))&((esk4_2(X18,X19)=u1_struct_0(X19)|v1_tsep_1(X19,X18)|~m1_pre_topc(X19,X18)|~l1_pre_topc(X18))&(~v3_pre_topc(esk4_2(X18,X19),X18)|v1_tsep_1(X19,X18)|~m1_pre_topc(X19,X18)|~l1_pre_topc(X18))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tsep_1])])])])])).
cnf(c_0_15, plain, (v3_pre_topc(u1_struct_0(X1),X2)|v2_struct_0(X2)|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k6_tmap_1(X2,u1_struct_0(X1))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_16, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))=k8_tmap_1(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20, plain, (X1=k6_tmap_1(X2,X4)|v2_struct_0(X2)|X1!=k8_tmap_1(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|X4!=u1_struct_0(X3)|~v1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, plain, (v2_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22, plain, (l1_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23, plain, (v1_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_24, plain, (v1_tsep_1(X2,X1)|~v3_pre_topc(esk4_2(X1,X2),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_25, plain, (esk4_2(X1,X2)=u1_struct_0(X2)|v1_tsep_1(X2,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_26, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|v3_pre_topc(u1_struct_0(X1),esk1_0)|k6_tmap_1(esk1_0,u1_struct_0(X1))!=k8_tmap_1(esk1_0,esk2_0)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_18])]), c_0_19])).
cnf(c_0_27, plain, (k6_tmap_1(X1,u1_struct_0(X2))=k8_tmap_1(X1,X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_20])]), c_0_21]), c_0_22]), c_0_10]), c_0_23])).
cnf(c_0_28, plain, (v1_tsep_1(X1,X2)|~m1_pre_topc(X1,X2)|~v3_pre_topc(u1_struct_0(X1),X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|v3_pre_topc(u1_struct_0(X1),esk1_0)|k8_tmap_1(esk1_0,X1)!=k8_tmap_1(esk1_0,esk2_0)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_17]), c_0_18])]), c_0_19])).
cnf(c_0_30, negated_conjecture, (~v1_tsep_1(esk2_0,esk1_0)|~m1_pre_topc(esk2_0,esk1_0)|g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))!=k8_tmap_1(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_31, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_32, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|v1_tsep_1(X1,esk1_0)|k8_tmap_1(esk1_0,X1)!=k8_tmap_1(esk1_0,esk2_0)|~m1_pre_topc(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_17])])).
cnf(c_0_33, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))!=k8_tmap_1(esk1_0,esk2_0)|~v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_31])])).
cnf(c_0_34, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_32, c_0_31])).
cnf(c_0_35, plain, (g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))=k6_tmap_1(X2,X1)|v2_struct_0(X2)|~v3_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_36, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))!=k8_tmap_1(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_34])])).
cnf(c_0_37, plain, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=k6_tmap_1(X1,u1_struct_0(X2))|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~v3_pre_topc(u1_struct_0(X2),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_35, c_0_10])).
cnf(c_0_38, negated_conjecture, (k6_tmap_1(esk1_0,u1_struct_0(X1))!=k8_tmap_1(esk1_0,esk2_0)|~m1_pre_topc(X1,esk1_0)|~v3_pre_topc(u1_struct_0(X1),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_17]), c_0_18])]), c_0_19])).
cnf(c_0_39, plain, (v3_pre_topc(X3,X2)|~v1_tsep_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|X3!=u1_struct_0(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_40, negated_conjecture, (k8_tmap_1(esk1_0,X1)!=k8_tmap_1(esk1_0,esk2_0)|~m1_pre_topc(X1,esk1_0)|~v3_pre_topc(u1_struct_0(X1),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_27]), c_0_17]), c_0_18])]), c_0_19])).
cnf(c_0_41, plain, (v3_pre_topc(u1_struct_0(X1),X2)|~v1_tsep_1(X1,X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_39]), c_0_10])).
cnf(c_0_42, negated_conjecture, (k8_tmap_1(esk1_0,X1)!=k8_tmap_1(esk1_0,esk2_0)|~v1_tsep_1(X1,esk1_0)|~m1_pre_topc(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_17])])).
cnf(c_0_43, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_34]), c_0_31])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 44
# Proof object clause steps            : 31
# Proof object formula steps           : 13
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 34
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 22
# Processed clauses                    : 38
# ...of these trivial                  : 1
# ...subsumed                          : 2
# ...remaining for further processing  : 35
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 5
# Generated clauses                    : 30
# ...of the previous two non-trivial   : 26
# Contextual simplify-reflections      : 5
# Paramodulations                      : 27
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 28
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 21
# Current number of unprocessed clauses: 10
# ...number of literals in the above   : 81
# Current number of archived formulas  : 0
# Current number of archived clauses   : 5
# Clause-clause subsumption calls (NU) : 119
# Rec. Clause-clause subsumption calls : 8
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2732

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.003 s
# Total time               : 0.009 s
# Maximum resident set size: 28796 pages
