# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/0b9e44daf6bd4dbe8d4e23d100db69fa.E_input', d1_tarski)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/0b9e44daf6bd4dbe8d4e23d100db69fa.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/0b9e44daf6bd4dbe8d4e23d100db69fa.E_input', t70_enumset1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/0b9e44daf6bd4dbe8d4e23d100db69fa.E_input', t113_zfmisc_1)).
fof(t73_zfmisc_1, axiom, ![X1, X2, X3]:(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_xboole_0<=>(r2_hidden(X1,X3)&r2_hidden(X2,X3))), file('data/MPTP/proofs/0b9e44daf6bd4dbe8d4e23d100db69fa.E_input', t73_zfmisc_1)).
fof(t152_zfmisc_1, axiom, ![X1, X2]:~(r2_hidden(X1,k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/0b9e44daf6bd4dbe8d4e23d100db69fa.E_input', t152_zfmisc_1)).
fof(t48_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>X4=k3_mcart_1(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4),k7_mcart_1(X1,X2,X3,X4)))))))), file('data/MPTP/proofs/0b9e44daf6bd4dbe8d4e23d100db69fa.E_input', t48_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/0b9e44daf6bd4dbe8d4e23d100db69fa.E_input', d3_zfmisc_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/0b9e44daf6bd4dbe8d4e23d100db69fa.E_input', d3_mcart_1)).
fof(t51_mcart_1, conjecture, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(~(X4=k5_mcart_1(X1,X2,X3,X4))&(~(X4=k6_mcart_1(X1,X2,X3,X4))&~(X4=k7_mcart_1(X1,X2,X3,X4)))))))))), file('data/MPTP/proofs/0b9e44daf6bd4dbe8d4e23d100db69fa.E_input', t51_mcart_1)).
fof(t29_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4, X5]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k3_mcart_1(X3,X4,X5))))))), file('data/MPTP/proofs/0b9e44daf6bd4dbe8d4e23d100db69fa.E_input', t29_mcart_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/0b9e44daf6bd4dbe8d4e23d100db69fa.E_input', t3_boole)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/0b9e44daf6bd4dbe8d4e23d100db69fa.E_input', t7_mcart_1)).
fof(t20_mcart_1, axiom, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1)))), file('data/MPTP/proofs/0b9e44daf6bd4dbe8d4e23d100db69fa.E_input', t20_mcart_1)).
fof(d1_enumset1, axiom, ![X1, X2, X3, X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3)))))), file('data/MPTP/proofs/0b9e44daf6bd4dbe8d4e23d100db69fa.E_input', d1_enumset1)).
fof(c_0_15, plain, ![X6, X7, X8, X9, X10, X11]:(((~r2_hidden(X8,X7)|X8=X6|X7!=k1_tarski(X6))&(X9!=X6|r2_hidden(X9,X7)|X7!=k1_tarski(X6)))&((~r2_hidden(esk1_2(X10,X11),X11)|esk1_2(X10,X11)!=X10|X11=k1_tarski(X10))&(r2_hidden(esk1_2(X10,X11),X11)|esk1_2(X10,X11)=X10|X11=k1_tarski(X10)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_16, plain, ![X13]:k2_tarski(X13,X13)=k1_tarski(X13), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_17, plain, ![X14, X15]:k1_enumset1(X14,X14,X15)=k2_tarski(X14,X15), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_18, plain, ![X16, X17]:((k2_zfmisc_1(X16,X17)!=k1_xboole_0|(X16=k1_xboole_0|X17=k1_xboole_0))&((X16!=k1_xboole_0|k2_zfmisc_1(X16,X17)=k1_xboole_0)&(X17!=k1_xboole_0|k2_zfmisc_1(X16,X17)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_19, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_22, plain, ![X18, X19, X20]:(((r2_hidden(X18,X20)|k4_xboole_0(k2_tarski(X18,X19),X20)!=k1_xboole_0)&(r2_hidden(X19,X20)|k4_xboole_0(k2_tarski(X18,X19),X20)!=k1_xboole_0))&(~r2_hidden(X18,X20)|~r2_hidden(X19,X20)|k4_xboole_0(k2_tarski(X18,X19),X20)=k1_xboole_0)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t73_zfmisc_1])])])).
fof(c_0_23, plain, ![X21, X22]:~r2_hidden(X21,k2_zfmisc_1(X21,X22)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_zfmisc_1])])).
cnf(c_0_24, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_25, plain, ![X23, X24, X25, X26]:(X23=k1_xboole_0|(X24=k1_xboole_0|(X25=k1_xboole_0|(~m1_subset_1(X26,k3_zfmisc_1(X23,X24,X25))|X26=k3_mcart_1(k5_mcart_1(X23,X24,X25,X26),k6_mcart_1(X23,X24,X25,X26),k7_mcart_1(X23,X24,X25,X26)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t48_mcart_1])])])])).
fof(c_0_26, plain, ![X27, X28, X29]:k3_zfmisc_1(X27,X28,X29)=k2_zfmisc_1(k2_zfmisc_1(X27,X28),X29), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_27, plain, ![X30, X31, X32]:k3_mcart_1(X30,X31,X32)=k4_tarski(k4_tarski(X30,X31),X32), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_28, negated_conjecture, ~(![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(~(X4=k5_mcart_1(X1,X2,X3,X4))&(~(X4=k6_mcart_1(X1,X2,X3,X4))&~(X4=k7_mcart_1(X1,X2,X3,X4))))))))))), inference(assume_negation,[status(cth)],[t51_mcart_1])).
cnf(c_0_29, plain, (X1=X3|X2!=k1_enumset1(X3,X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_20]), c_0_21])).
fof(c_0_30, plain, ![X37, X39, X40, X41]:((r2_hidden(esk6_1(X37),X37)|X37=k1_xboole_0)&((~r2_hidden(X39,X37)|esk6_1(X37)!=k3_mcart_1(X39,X40,X41)|X37=k1_xboole_0)&(~r2_hidden(X40,X37)|esk6_1(X37)!=k3_mcart_1(X39,X40,X41)|X37=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t29_mcart_1])])])])])])])).
cnf(c_0_31, plain, (r2_hidden(X1,X2)|k4_xboole_0(k2_tarski(X3,X1),X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_32, plain, ![X42]:k4_xboole_0(X42,k1_xboole_0)=X42, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_33, plain, (~r2_hidden(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_34, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_24])).
fof(c_0_35, plain, ![X43, X44]:(k1_mcart_1(k4_tarski(X43,X44))=X43&k2_mcart_1(k4_tarski(X43,X44))=X44), inference(variable_rename,[status(thm)],[t7_mcart_1])).
cnf(c_0_36, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k3_mcart_1(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4),k7_mcart_1(X1,X2,X3,X4))|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_37, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_38, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_39, negated_conjecture, (esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&(esk4_0!=k1_xboole_0&(m1_subset_1(esk5_0,k3_zfmisc_1(esk2_0,esk3_0,esk4_0))&(esk5_0=k5_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)|(esk5_0=k6_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)|esk5_0=k7_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_28])])])])).
cnf(c_0_40, plain, (X1=X2|~r2_hidden(X1,k1_enumset1(X2,X2,X2))), inference(er,[status(thm)],[c_0_29])).
cnf(c_0_41, plain, (r2_hidden(esk6_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_42, plain, (r2_hidden(X1,X2)|k4_xboole_0(k1_enumset1(X3,X3,X1),X2)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_31, c_0_21])).
cnf(c_0_43, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_44, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_45, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_46, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X4=k4_tarski(k4_tarski(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4)),k7_mcart_1(X1,X2,X3,X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_47, negated_conjecture, (m1_subset_1(esk5_0,k3_zfmisc_1(esk2_0,esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_48, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk6_1(X2)!=k3_mcart_1(X3,X1,X4)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_49, plain, (esk6_1(k1_enumset1(X1,X1,X1))=X1|k1_enumset1(X1,X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_50, plain, (k1_enumset1(X1,X1,X2)!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])).
cnf(c_0_51, plain, (k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_52, negated_conjecture, (m1_subset_1(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0))), inference(rw,[status(thm)],[c_0_47, c_0_37])).
cnf(c_0_53, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_54, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_55, negated_conjecture, (esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_56, plain, ![X45, X46, X47]:((X45!=k1_mcart_1(X45)|X45!=k4_tarski(X46,X47))&(X45!=k2_mcart_1(X45)|X45!=k4_tarski(X46,X47))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t20_mcart_1])])])])])).
cnf(c_0_57, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk6_1(X2)!=k3_mcart_1(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_58, plain, (X2=k1_xboole_0|esk6_1(X2)!=k4_tarski(k4_tarski(X3,X1),X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_48, c_0_38])).
cnf(c_0_59, plain, (esk6_1(k1_enumset1(X1,X1,X1))=X1), inference(sr,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_60, negated_conjecture, (k7_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)=k2_mcart_1(esk5_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53]), c_0_54]), c_0_55])).
cnf(c_0_61, plain, (X1!=k2_mcart_1(X1)|X1!=k4_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_56])).
fof(c_0_62, plain, ![X48, X49, X50, X51, X52, X53, X54, X55, X56, X57]:(((~r2_hidden(X52,X51)|(X52=X48|(X52=X49|X52=X50))|X51!=k1_enumset1(X48,X49,X50))&((X53!=X48|r2_hidden(X53,X51)|X51!=k1_enumset1(X48,X49,X50))&((X53!=X49|r2_hidden(X53,X51)|X51!=k1_enumset1(X48,X49,X50))&(X53!=X50|r2_hidden(X53,X51)|X51!=k1_enumset1(X48,X49,X50)))))&(((esk7_4(X54,X55,X56,X57)!=X54|~r2_hidden(esk7_4(X54,X55,X56,X57),X57)|X57=k1_enumset1(X54,X55,X56))&((esk7_4(X54,X55,X56,X57)!=X55|~r2_hidden(esk7_4(X54,X55,X56,X57),X57)|X57=k1_enumset1(X54,X55,X56))&(esk7_4(X54,X55,X56,X57)!=X56|~r2_hidden(esk7_4(X54,X55,X56,X57),X57)|X57=k1_enumset1(X54,X55,X56))))&(r2_hidden(esk7_4(X54,X55,X56,X57),X57)|(esk7_4(X54,X55,X56,X57)=X54|(esk7_4(X54,X55,X56,X57)=X55|esk7_4(X54,X55,X56,X57)=X56))|X57=k1_enumset1(X54,X55,X56)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])).
cnf(c_0_63, plain, (X2=k1_xboole_0|esk6_1(X2)!=k4_tarski(k4_tarski(X1,X3),X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_57, c_0_38])).
cnf(c_0_64, plain, (~r2_hidden(X1,k1_enumset1(k4_tarski(k4_tarski(X2,X1),X3),k4_tarski(k4_tarski(X2,X1),X3),k4_tarski(k4_tarski(X2,X1),X3)))), inference(er,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_50])])).
cnf(c_0_65, negated_conjecture, (k4_tarski(k4_tarski(k5_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0),k6_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)),k2_mcart_1(esk5_0))=esk5_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_60]), c_0_52])]), c_0_55]), c_0_54]), c_0_53])).
cnf(c_0_66, negated_conjecture, (esk5_0=k5_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)|esk5_0=k6_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)|esk5_0=k7_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_67, plain, (k4_tarski(X1,X2)!=X2), inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_61]), c_0_45])).
cnf(c_0_68, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X2,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_69, plain, (~r2_hidden(X1,k1_enumset1(k4_tarski(k4_tarski(X1,X2),X3),k4_tarski(k4_tarski(X1,X2),X3),k4_tarski(k4_tarski(X1,X2),X3)))), inference(er,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_59]), c_0_50])])).
cnf(c_0_70, negated_conjecture, (~r2_hidden(k6_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0),k1_enumset1(esk5_0,esk5_0,esk5_0))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_71, negated_conjecture, (k6_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)=esk5_0|k5_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)=esk5_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_66]), c_0_52])]), c_0_67]), c_0_55]), c_0_54]), c_0_53])).
cnf(c_0_72, plain, (r2_hidden(X1,k1_enumset1(X1,X2,X3))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_68])])).
cnf(c_0_73, negated_conjecture, (~r2_hidden(k5_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0),k1_enumset1(esk5_0,esk5_0,esk5_0))), inference(spm,[status(thm)],[c_0_69, c_0_65])).
cnf(c_0_74, negated_conjecture, (k5_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_72])])).
cnf(c_0_75, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_74]), c_0_72])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 76
# Proof object clause steps            : 45
# Proof object formula steps           : 31
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 39
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 37
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 33
# Processed clauses                    : 143
# ...of these trivial                  : 2
# ...subsumed                          : 25
# ...remaining for further processing  : 116
# Other redundant clauses eliminated   : 73
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 5
# Generated clauses                    : 673
# ...of the previous two non-trivial   : 532
# Contextual simplify-reflections      : 0
# Paramodulations                      : 597
# Factorizations                       : 6
# Equation resolutions                 : 73
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 65
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 13
#    Non-unit-clauses                  : 39
# Current number of unprocessed clauses: 450
# ...number of literals in the above   : 2527
# Current number of archived formulas  : 0
# Current number of archived clauses   : 45
# Clause-clause subsumption calls (NU) : 218
# Rec. Clause-clause subsumption calls : 102
# Non-unit clause-clause subsumptions  : 21
# Unit Clause-clause subsumption calls : 18
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 11302

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.006 s
# Total time               : 0.018 s
# Maximum resident set size: 28796 pages
