# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t36_tex_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>v2_tex_2(k6_domain_1(u1_struct_0(X1),X2),X1))), file('data/MPTP/proofs/0c15e0525659486a95672163621bf960.E_input', t36_tex_2)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/0c15e0525659486a95672163621bf960.E_input', dt_l1_pre_topc)).
fof(t63_subset_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))), file('data/MPTP/proofs/0c15e0525659486a95672163621bf960.E_input', t63_subset_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/0c15e0525659486a95672163621bf960.E_input', t69_enumset1)).
fof(fc2_xboole_0, axiom, ![X1]:~(v1_xboole_0(k1_tarski(X1))), file('data/MPTP/proofs/0c15e0525659486a95672163621bf960.E_input', fc2_xboole_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/0c15e0525659486a95672163621bf960.E_input', fc2_struct_0)).
fof(t10_tsep_1, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>?[X3]:((~(v2_struct_0(X3))&(v1_pre_topc(X3)&m1_pre_topc(X3,X1)))&X2=u1_struct_0(X3)))), file('data/MPTP/proofs/0c15e0525659486a95672163621bf960.E_input', t10_tsep_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/0c15e0525659486a95672163621bf960.E_input', t2_subset)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/0c15e0525659486a95672163621bf960.E_input', redefinition_k6_domain_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/0c15e0525659486a95672163621bf960.E_input', d1_tarski)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/0c15e0525659486a95672163621bf960.E_input', dt_m1_pre_topc)).
fof(t1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/0c15e0525659486a95672163621bf960.E_input', t1_tsep_1)).
fof(d1_tex_2, axiom, ![X1]:(~(v1_xboole_0(X1))=>(v1_zfmisc_1(X1)<=>?[X2]:(m1_subset_1(X2,X1)&X1=k6_domain_1(X1,X2)))), file('data/MPTP/proofs/0c15e0525659486a95672163621bf960.E_input', d1_tex_2)).
fof(t2_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>m1_pre_topc(X1,X1)), file('data/MPTP/proofs/0c15e0525659486a95672163621bf960.E_input', t2_tsep_1)).
fof(cc1_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>v2_pre_topc(X2))), file('data/MPTP/proofs/0c15e0525659486a95672163621bf960.E_input', cc1_pre_topc)).
fof(cc25_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>((~(v2_struct_0(X2))&(v7_struct_0(X2)&v2_pre_topc(X2)))=>(~(v2_struct_0(X2))&(v1_tdlat_3(X2)&v2_tdlat_3(X2)))))), file('data/MPTP/proofs/0c15e0525659486a95672163621bf960.E_input', cc25_tex_2)).
fof(d1_tex_1, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>(v7_struct_0(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&u1_struct_0(X1)=k6_domain_1(u1_struct_0(X1),X2)))), file('data/MPTP/proofs/0c15e0525659486a95672163621bf960.E_input', d1_tex_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/0c15e0525659486a95672163621bf960.E_input', t1_subset)).
fof(t26_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>(v2_tex_2(X3,X1)<=>v1_tdlat_3(X2)))))), file('data/MPTP/proofs/0c15e0525659486a95672163621bf960.E_input', t26_tex_2)).
fof(c_0_19, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>v2_tex_2(k6_domain_1(u1_struct_0(X1),X2),X1)))), inference(assume_negation,[status(cth)],[t36_tex_2])).
fof(c_0_20, plain, ![X6]:(~l1_pre_topc(X6)|l1_struct_0(X6)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_21, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&~v2_tex_2(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_19])])])])).
fof(c_0_22, plain, ![X7, X8]:(~r2_hidden(X7,X8)|m1_subset_1(k1_tarski(X7),k1_zfmisc_1(X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_subset_1])])).
fof(c_0_23, plain, ![X9]:k2_tarski(X9,X9)=k1_tarski(X9), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_24, plain, ![X10]:~v1_xboole_0(k1_tarski(X10)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_xboole_0])])).
fof(c_0_25, plain, ![X11]:(v2_struct_0(X11)|~l1_struct_0(X11)|~v1_xboole_0(u1_struct_0(X11))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_26, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_28, plain, ![X12, X13]:(((~v2_struct_0(esk3_2(X12,X13))|(v1_xboole_0(X13)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12))))|(v2_struct_0(X12)|~l1_pre_topc(X12)))&((v1_pre_topc(esk3_2(X12,X13))|(v1_xboole_0(X13)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12))))|(v2_struct_0(X12)|~l1_pre_topc(X12)))&(m1_pre_topc(esk3_2(X12,X13),X12)|(v1_xboole_0(X13)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12))))|(v2_struct_0(X12)|~l1_pre_topc(X12)))))&(X13=u1_struct_0(esk3_2(X12,X13))|(v1_xboole_0(X13)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12))))|(v2_struct_0(X12)|~l1_pre_topc(X12)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t10_tsep_1])])])])])])).
cnf(c_0_29, plain, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31, plain, (~v1_xboole_0(k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_32, plain, ![X15, X16]:(~m1_subset_1(X15,X16)|(v1_xboole_0(X16)|r2_hidden(X15,X16))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_33, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_34, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, negated_conjecture, (l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_36, plain, (m1_pre_topc(esk3_2(X1,X2),X1)|v1_xboole_0(X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, plain, (m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(X2))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_38, plain, (~v1_xboole_0(k2_tarski(X1,X1))), inference(rw,[status(thm)],[c_0_31, c_0_30])).
cnf(c_0_39, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_41, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])])).
fof(c_0_42, plain, ![X17, X18]:(v1_xboole_0(X17)|~m1_subset_1(X18,X17)|k6_domain_1(X17,X18)=k1_tarski(X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_43, plain, ![X19, X20, X21, X22, X23, X24]:(((~r2_hidden(X21,X20)|X21=X19|X20!=k1_tarski(X19))&(X22!=X19|r2_hidden(X22,X20)|X20!=k1_tarski(X19)))&((~r2_hidden(esk4_2(X23,X24),X24)|esk4_2(X23,X24)!=X23|X24=k1_tarski(X23))&(r2_hidden(esk4_2(X23,X24),X24)|esk4_2(X23,X24)=X23|X24=k1_tarski(X23)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_44, plain, ![X26, X27]:(~l1_pre_topc(X26)|(~m1_pre_topc(X27,X26)|l1_pre_topc(X27))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_45, plain, (m1_pre_topc(esk3_2(X1,k2_tarski(X2,X2)),X1)|v2_struct_0(X1)|~r2_hidden(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])).
fof(c_0_47, plain, ![X28, X29]:(~l1_pre_topc(X28)|(~m1_pre_topc(X29,X28)|m1_subset_1(u1_struct_0(X29),k1_zfmisc_1(u1_struct_0(X28))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])).
cnf(c_0_48, plain, (X1=u1_struct_0(esk3_2(X2,X1))|v1_xboole_0(X1)|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_49, plain, ![X30, X32]:(((m1_subset_1(esk5_1(X30),X30)|~v1_zfmisc_1(X30)|v1_xboole_0(X30))&(X30=k6_domain_1(X30,esk5_1(X30))|~v1_zfmisc_1(X30)|v1_xboole_0(X30)))&(~m1_subset_1(X32,X30)|X30!=k6_domain_1(X30,X32)|v1_zfmisc_1(X30)|v1_xboole_0(X30))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_tex_2])])])])])])).
cnf(c_0_50, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_51, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_52, plain, ![X33]:(~l1_pre_topc(X33)|m1_pre_topc(X33,X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tsep_1])])).
cnf(c_0_53, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_54, negated_conjecture, (m1_pre_topc(esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0)),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_27])]), c_0_33])).
fof(c_0_55, plain, ![X34, X35]:(~v2_pre_topc(X34)|~l1_pre_topc(X34)|(~m1_pre_topc(X35,X34)|v2_pre_topc(X35))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_pre_topc])])])).
cnf(c_0_56, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_57, plain, (u1_struct_0(esk3_2(X1,k2_tarski(X2,X2)))=k2_tarski(X2,X2)|v2_struct_0(X1)|~r2_hidden(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_37]), c_0_38])).
cnf(c_0_58, plain, (v1_zfmisc_1(X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)|X2!=k6_domain_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_59, plain, (k2_tarski(X2,X2)=k6_domain_1(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_50, c_0_30])).
cnf(c_0_60, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X2)), inference(rw,[status(thm)],[c_0_51, c_0_30])).
fof(c_0_61, plain, ![X36, X37]:((~v2_struct_0(X37)|(v2_struct_0(X37)|(~v7_struct_0(X37)|~v2_pre_topc(X37)))|~m1_pre_topc(X37,X36)|(v2_struct_0(X36)|~l1_pre_topc(X36)))&((v1_tdlat_3(X37)|(v2_struct_0(X37)|(~v7_struct_0(X37)|~v2_pre_topc(X37)))|~m1_pre_topc(X37,X36)|(v2_struct_0(X36)|~l1_pre_topc(X36)))&(v2_tdlat_3(X37)|(v2_struct_0(X37)|(~v7_struct_0(X37)|~v2_pre_topc(X37)))|~m1_pre_topc(X37,X36)|(v2_struct_0(X36)|~l1_pre_topc(X36))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc25_tex_2])])])])])).
cnf(c_0_62, plain, (m1_pre_topc(X1,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_63, negated_conjecture, (l1_pre_topc(esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_27])])).
cnf(c_0_64, plain, (v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_65, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_66, plain, ![X38, X40]:(((m1_subset_1(esk6_1(X38),u1_struct_0(X38))|~v7_struct_0(X38)|(v2_struct_0(X38)|~l1_struct_0(X38)))&(u1_struct_0(X38)=k6_domain_1(u1_struct_0(X38),esk6_1(X38))|~v7_struct_0(X38)|(v2_struct_0(X38)|~l1_struct_0(X38))))&(~m1_subset_1(X40,u1_struct_0(X38))|u1_struct_0(X38)!=k6_domain_1(u1_struct_0(X38),X40)|v7_struct_0(X38)|(v2_struct_0(X38)|~l1_struct_0(X38)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_tex_1])])])])])])).
cnf(c_0_67, plain, (m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X3)|~m1_pre_topc(esk3_2(X3,k2_tarski(X1,X1)),X2)|~r2_hidden(X1,u1_struct_0(X3))|~l1_pre_topc(X2)|~l1_pre_topc(X3)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_68, plain, (v1_zfmisc_1(X1)|v1_xboole_0(X1)|k2_tarski(X2,X2)!=X1|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
fof(c_0_69, plain, ![X41, X42]:(~r2_hidden(X41,X42)|m1_subset_1(X41,X42)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_70, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X1,X1)), inference(er,[status(thm)],[c_0_60])).
cnf(c_0_71, plain, (v1_xboole_0(X2)|v2_struct_0(X1)|~v2_struct_0(esk3_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_72, plain, (v1_tdlat_3(X1)|v2_struct_0(X1)|v2_struct_0(X2)|~v7_struct_0(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_73, negated_conjecture, (m1_pre_topc(esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0)),esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0)))), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_74, negated_conjecture, (v2_pre_topc(esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_54]), c_0_27]), c_0_65])])).
cnf(c_0_75, plain, (v7_struct_0(X2)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|u1_struct_0(X2)!=k6_domain_1(u1_struct_0(X2),X1)|~l1_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_76, plain, (X1=k6_domain_1(X1,esk5_1(X1))|v1_xboole_0(X1)|~v1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_77, plain, (m1_subset_1(esk5_1(X1),X1)|v1_xboole_0(X1)|~v1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_78, negated_conjecture, (m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_54]), c_0_46]), c_0_27])]), c_0_33])).
cnf(c_0_79, plain, (v1_zfmisc_1(k2_tarski(X1,X1))|~m1_subset_1(X1,k2_tarski(X1,X1))), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_68]), c_0_38])).
cnf(c_0_80, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_81, plain, (r2_hidden(X1,k2_tarski(X1,X1))), inference(er,[status(thm)],[c_0_70])).
fof(c_0_82, plain, ![X43, X44, X45]:((~v2_tex_2(X45,X43)|v1_tdlat_3(X44)|X45!=u1_struct_0(X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X43)))|(v2_struct_0(X44)|~m1_pre_topc(X44,X43))|(v2_struct_0(X43)|~l1_pre_topc(X43)))&(~v1_tdlat_3(X44)|v2_tex_2(X45,X43)|X45!=u1_struct_0(X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X43)))|(v2_struct_0(X44)|~m1_pre_topc(X44,X43))|(v2_struct_0(X43)|~l1_pre_topc(X43)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t26_tex_2])])])])])).
cnf(c_0_83, negated_conjecture, (~v2_tex_2(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_84, plain, (v2_struct_0(X1)|~r2_hidden(X2,u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_struct_0(esk3_2(X1,k2_tarski(X2,X2)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_37]), c_0_38])).
cnf(c_0_85, negated_conjecture, (v1_tdlat_3(esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0)))|v2_struct_0(esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0)))|~v7_struct_0(esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_63]), c_0_74])])).
cnf(c_0_86, plain, (v7_struct_0(X1)|v2_struct_0(X1)|~v1_zfmisc_1(u1_struct_0(X1))|~l1_struct_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77]), c_0_34])).
cnf(c_0_87, negated_conjecture, (u1_struct_0(esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0)))=k2_tarski(esk2_0,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_78]), c_0_27])]), c_0_38]), c_0_33])).
cnf(c_0_88, plain, (v1_zfmisc_1(k2_tarski(X1,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_81])])).
cnf(c_0_89, negated_conjecture, (l1_struct_0(esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0)))), inference(spm,[status(thm)],[c_0_26, c_0_63])).
cnf(c_0_90, negated_conjecture, (~v2_struct_0(esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_78]), c_0_27])]), c_0_38]), c_0_33])).
cnf(c_0_91, plain, (v2_tex_2(X2,X3)|v2_struct_0(X1)|v2_struct_0(X3)|~v1_tdlat_3(X1)|X2!=u1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_pre_topc(X1,X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_92, negated_conjecture, (~v2_tex_2(k2_tarski(esk2_0,esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_59]), c_0_40])]), c_0_41])).
cnf(c_0_93, negated_conjecture, (v1_tdlat_3(esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0)))|~v7_struct_0(esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_46]), c_0_27])]), c_0_33])).
cnf(c_0_94, negated_conjecture, (v7_struct_0(esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_88]), c_0_89])]), c_0_90])).
cnf(c_0_95, negated_conjecture, (v2_struct_0(X1)|u1_struct_0(X1)!=k2_tarski(esk2_0,esk2_0)|~v1_tdlat_3(X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_78]), c_0_27])]), c_0_92]), c_0_33])).
cnf(c_0_96, negated_conjecture, (v1_tdlat_3(esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_94])])).
cnf(c_0_97, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_95]), c_0_87]), c_0_96]), c_0_54])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 98
# Proof object clause steps            : 59
# Proof object formula steps           : 39
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 26
# Proof object simplifying inferences  : 60
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 36
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 34
# Processed clauses                    : 2084
# ...of these trivial                  : 6
# ...subsumed                          : 1443
# ...remaining for further processing  : 635
# Other redundant clauses eliminated   : 181
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 39
# Backward-rewritten                   : 9
# Generated clauses                    : 15387
# ...of the previous two non-trivial   : 14571
# Contextual simplify-reflections      : 79
# Paramodulations                      : 15153
# Factorizations                       : 14
# Equation resolutions                 : 219
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 551
#    Positive orientable unit clauses  : 23
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 521
# Current number of unprocessed clauses: 12400
# ...number of literals in the above   : 96077
# Current number of archived formulas  : 0
# Current number of archived clauses   : 84
# Clause-clause subsumption calls (NU) : 62655
# Rec. Clause-clause subsumption calls : 11005
# Non-unit clause-clause subsumptions  : 784
# Unit Clause-clause subsumption calls : 330
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 388506

# -------------------------------------------------
# User time                : 0.250 s
# System time              : 0.016 s
# Total time               : 0.267 s
# Maximum resident set size: 28744 pages
