# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C45_F1_PI_AE_Q4_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t45_mcart_1, conjecture, ![X1, X2, X3, X4, X5]:k3_zfmisc_1(k1_tarski(X1),k2_tarski(X2,X3),k2_tarski(X4,X5))=k2_enumset1(k3_mcart_1(X1,X2,X4),k3_mcart_1(X1,X3,X4),k3_mcart_1(X1,X2,X5),k3_mcart_1(X1,X3,X5)), file('data/MPTP/proofs/0c30d8994de84a3daf483a6249fa77f6.E_input', t45_mcart_1)).
fof(t36_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k1_tarski(X1),k2_tarski(X2,X3))=k2_tarski(k4_tarski(X1,X2),k4_tarski(X1,X3))&k2_zfmisc_1(k2_tarski(X1,X2),k1_tarski(X3))=k2_tarski(k4_tarski(X1,X3),k4_tarski(X2,X3))), file('data/MPTP/proofs/0c30d8994de84a3daf483a6249fa77f6.E_input', t36_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/0c30d8994de84a3daf483a6249fa77f6.E_input', t69_enumset1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/0c30d8994de84a3daf483a6249fa77f6.E_input', d3_zfmisc_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/0c30d8994de84a3daf483a6249fa77f6.E_input', d3_mcart_1)).
fof(l53_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X4)), file('data/MPTP/proofs/0c30d8994de84a3daf483a6249fa77f6.E_input', l53_enumset1)).
fof(t120_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&k2_zfmisc_1(X3,k2_xboole_0(X1,X2))=k2_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))), file('data/MPTP/proofs/0c30d8994de84a3daf483a6249fa77f6.E_input', t120_zfmisc_1)).
fof(t41_enumset1, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_xboole_0(k1_tarski(X1),k1_tarski(X2)), file('data/MPTP/proofs/0c30d8994de84a3daf483a6249fa77f6.E_input', t41_enumset1)).
fof(c_0_8, negated_conjecture, ~(![X1, X2, X3, X4, X5]:k3_zfmisc_1(k1_tarski(X1),k2_tarski(X2,X3),k2_tarski(X4,X5))=k2_enumset1(k3_mcart_1(X1,X2,X4),k3_mcart_1(X1,X3,X4),k3_mcart_1(X1,X2,X5),k3_mcart_1(X1,X3,X5))), inference(assume_negation,[status(cth)],[t45_mcart_1])).
fof(c_0_9, plain, ![X11, X12, X13]:(k2_zfmisc_1(k1_tarski(X11),k2_tarski(X12,X13))=k2_tarski(k4_tarski(X11,X12),k4_tarski(X11,X13))&k2_zfmisc_1(k2_tarski(X11,X12),k1_tarski(X13))=k2_tarski(k4_tarski(X11,X13),k4_tarski(X12,X13))), inference(variable_rename,[status(thm)],[t36_zfmisc_1])).
fof(c_0_10, plain, ![X14]:k2_tarski(X14,X14)=k1_tarski(X14), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_11, negated_conjecture, k3_zfmisc_1(k1_tarski(esk1_0),k2_tarski(esk2_0,esk3_0),k2_tarski(esk4_0,esk5_0))!=k2_enumset1(k3_mcart_1(esk1_0,esk2_0,esk4_0),k3_mcart_1(esk1_0,esk3_0,esk4_0),k3_mcart_1(esk1_0,esk2_0,esk5_0),k3_mcart_1(esk1_0,esk3_0,esk5_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
fof(c_0_12, plain, ![X15, X16, X17]:k3_zfmisc_1(X15,X16,X17)=k2_zfmisc_1(k2_zfmisc_1(X15,X16),X17), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_13, plain, ![X18, X19, X20]:k3_mcart_1(X18,X19,X20)=k4_tarski(k4_tarski(X18,X19),X20), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_14, plain, ![X21, X22, X23, X24]:k2_enumset1(X21,X22,X23,X24)=k2_xboole_0(k2_tarski(X21,X22),k2_tarski(X23,X24)), inference(variable_rename,[status(thm)],[l53_enumset1])).
fof(c_0_15, plain, ![X25, X26, X27]:(k2_zfmisc_1(k2_xboole_0(X25,X26),X27)=k2_xboole_0(k2_zfmisc_1(X25,X27),k2_zfmisc_1(X26,X27))&k2_zfmisc_1(X27,k2_xboole_0(X25,X26))=k2_xboole_0(k2_zfmisc_1(X27,X25),k2_zfmisc_1(X27,X26))), inference(variable_rename,[status(thm)],[t120_zfmisc_1])).
cnf(c_0_16, plain, (k2_zfmisc_1(k2_tarski(X1,X2),k1_tarski(X3))=k2_tarski(k4_tarski(X1,X3),k4_tarski(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_18, plain, ![X28, X29]:k2_tarski(X28,X29)=k2_xboole_0(k1_tarski(X28),k1_tarski(X29)), inference(variable_rename,[status(thm)],[t41_enumset1])).
cnf(c_0_19, negated_conjecture, (k3_zfmisc_1(k1_tarski(esk1_0),k2_tarski(esk2_0,esk3_0),k2_tarski(esk4_0,esk5_0))!=k2_enumset1(k3_mcart_1(esk1_0,esk2_0,esk4_0),k3_mcart_1(esk1_0,esk3_0,esk4_0),k3_mcart_1(esk1_0,esk2_0,esk5_0),k3_mcart_1(esk1_0,esk3_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22, plain, (k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X4))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23, plain, (k2_zfmisc_1(k1_tarski(X1),k2_tarski(X2,X3))=k2_tarski(k4_tarski(X1,X2),k4_tarski(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_24, plain, (k2_zfmisc_1(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25, plain, (k2_zfmisc_1(k2_tarski(X1,X2),k2_tarski(X3,X3))=k2_tarski(k4_tarski(X1,X3),k4_tarski(X2,X3))), inference(rw,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_26, plain, (k2_tarski(X1,X2)=k2_xboole_0(k1_tarski(X1),k1_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(k2_tarski(esk1_0,esk1_0),k2_tarski(esk2_0,esk3_0)),k2_tarski(esk4_0,esk5_0))!=k2_xboole_0(k2_tarski(k4_tarski(k4_tarski(esk1_0,esk2_0),esk4_0),k4_tarski(k4_tarski(esk1_0,esk3_0),esk4_0)),k2_tarski(k4_tarski(k4_tarski(esk1_0,esk2_0),esk5_0),k4_tarski(k4_tarski(esk1_0,esk3_0),esk5_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_17]), c_0_20]), c_0_21]), c_0_21]), c_0_21]), c_0_21]), c_0_22])).
cnf(c_0_28, plain, (k2_zfmisc_1(k2_tarski(X1,X1),k2_tarski(X2,X3))=k2_tarski(k4_tarski(X1,X2),k4_tarski(X1,X3))), inference(rw,[status(thm)],[c_0_23, c_0_17])).
cnf(c_0_29, plain, (k2_xboole_0(k2_tarski(k4_tarski(X1,X2),k4_tarski(X3,X2)),k2_zfmisc_1(k2_tarski(X1,X3),X4))=k2_zfmisc_1(k2_tarski(X1,X3),k2_xboole_0(k2_tarski(X2,X2),X4))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_30, plain, (k2_tarski(X1,X2)=k2_xboole_0(k2_tarski(X1,X1),k2_tarski(X2,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_17]), c_0_17])).
cnf(c_0_31, negated_conjecture, (k2_xboole_0(k2_tarski(k4_tarski(k4_tarski(esk1_0,esk2_0),esk4_0),k4_tarski(k4_tarski(esk1_0,esk3_0),esk4_0)),k2_tarski(k4_tarski(k4_tarski(esk1_0,esk2_0),esk5_0),k4_tarski(k4_tarski(esk1_0,esk3_0),esk5_0)))!=k2_zfmisc_1(k2_tarski(k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk3_0)),k2_tarski(esk4_0,esk5_0))), inference(rw,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_32, plain, (k2_xboole_0(k2_tarski(k4_tarski(X1,X2),k4_tarski(X3,X2)),k2_tarski(k4_tarski(X1,X4),k4_tarski(X3,X4)))=k2_zfmisc_1(k2_tarski(X1,X3),k2_tarski(X2,X4))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_25]), c_0_30])).
cnf(c_0_33, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_32])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 34
# Proof object clause steps            : 17
# Proof object formula steps           : 17
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 2
# Proof object simplifying inferences  : 15
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 10
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 6
# Processed clauses                    : 126
# ...of these trivial                  : 48
# ...subsumed                          : 14
# ...remaining for further processing  : 64
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 13
# Generated clauses                    : 2310
# ...of the previous two non-trivial   : 1720
# Contextual simplify-reflections      : 0
# Paramodulations                      : 2310
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 51
#    Positive orientable unit clauses  : 37
#    Positive unorientable unit clauses: 14
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 1586
# ...number of literals in the above   : 1586
# Current number of archived formulas  : 0
# Current number of archived clauses   : 17
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 49
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 74
# BW rewrite match successes           : 32
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 33326

# -------------------------------------------------
# User time                : 0.021 s
# System time              : 0.000 s
# Total time               : 0.021 s
# Maximum resident set size: 28696 pages
