# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t41_ordinal1, conjecture, ![X1]:(v3_ordinal1(X1)=>(v4_ordinal1(X1)<=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X2,X1)=>r2_hidden(k1_ordinal1(X2),X1))))), file('data/MPTP/proofs/0c45c0611536468bbc9f7014f489d076.E_input', t41_ordinal1)).
fof(t34_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,k1_ordinal1(X2))<=>r1_ordinal1(X1,X2)))), file('data/MPTP/proofs/0c45c0611536468bbc9f7014f489d076.E_input', t34_ordinal1)).
fof(d1_ordinal1, axiom, ![X1]:k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1)), file('data/MPTP/proofs/0c45c0611536468bbc9f7014f489d076.E_input', d1_ordinal1)).
fof(connectedness_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)|r1_ordinal1(X2,X1))), file('data/MPTP/proofs/0c45c0611536468bbc9f7014f489d076.E_input', connectedness_r1_ordinal1)).
fof(t38_ordinal1, axiom, ![X1]:~(![X2]:(v3_ordinal1(X2)=>r2_hidden(X2,X1))), file('data/MPTP/proofs/0c45c0611536468bbc9f7014f489d076.E_input', t38_ordinal1)).
fof(redefinition_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)<=>r1_tarski(X1,X2))), file('data/MPTP/proofs/0c45c0611536468bbc9f7014f489d076.E_input', redefinition_r1_ordinal1)).
fof(d2_ordinal1, axiom, ![X1]:(v1_ordinal1(X1)<=>![X2]:(r2_hidden(X2,X1)=>r1_tarski(X2,X1))), file('data/MPTP/proofs/0c45c0611536468bbc9f7014f489d076.E_input', d2_ordinal1)).
fof(cc1_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>(v1_ordinal1(X1)&v2_ordinal1(X1))), file('data/MPTP/proofs/0c45c0611536468bbc9f7014f489d076.E_input', cc1_ordinal1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/0c45c0611536468bbc9f7014f489d076.E_input', d3_tarski)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/0c45c0611536468bbc9f7014f489d076.E_input', t7_ordinal1)).
fof(d4_tarski, axiom, ![X1, X2]:(X2=k3_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X3,X4)&r2_hidden(X4,X1)))), file('data/MPTP/proofs/0c45c0611536468bbc9f7014f489d076.E_input', d4_tarski)).
fof(t23_ordinal1, axiom, ![X1, X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)=>v3_ordinal1(X1))), file('data/MPTP/proofs/0c45c0611536468bbc9f7014f489d076.E_input', t23_ordinal1)).
fof(t35_ordinal1, axiom, ![X1]:(![X2]:(r2_hidden(X2,X1)=>v3_ordinal1(X2))=>v3_ordinal1(k3_tarski(X1))), file('data/MPTP/proofs/0c45c0611536468bbc9f7014f489d076.E_input', t35_ordinal1)).
fof(t24_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>~((~(r2_hidden(X1,X2))&(~(X1=X2)&~(r2_hidden(X2,X1))))))), file('data/MPTP/proofs/0c45c0611536468bbc9f7014f489d076.E_input', t24_ordinal1)).
fof(t10_ordinal1, axiom, ![X1]:r2_hidden(X1,k1_ordinal1(X1)), file('data/MPTP/proofs/0c45c0611536468bbc9f7014f489d076.E_input', t10_ordinal1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/0c45c0611536468bbc9f7014f489d076.E_input', d10_xboole_0)).
fof(t33_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)<=>r1_ordinal1(k1_ordinal1(X1),X2)))), file('data/MPTP/proofs/0c45c0611536468bbc9f7014f489d076.E_input', t33_ordinal1)).
fof(t29_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>v3_ordinal1(k1_ordinal1(X1))), file('data/MPTP/proofs/0c45c0611536468bbc9f7014f489d076.E_input', t29_ordinal1)).
fof(d6_ordinal1, axiom, ![X1]:(v4_ordinal1(X1)<=>X1=k3_tarski(X1)), file('data/MPTP/proofs/0c45c0611536468bbc9f7014f489d076.E_input', d6_ordinal1)).
fof(c_0_19, negated_conjecture, ~(![X1]:(v3_ordinal1(X1)=>(v4_ordinal1(X1)<=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X2,X1)=>r2_hidden(k1_ordinal1(X2),X1)))))), inference(assume_negation,[status(cth)],[t41_ordinal1])).
fof(c_0_20, plain, ![X8, X9]:((~r2_hidden(X8,k1_ordinal1(X9))|r1_ordinal1(X8,X9)|~v3_ordinal1(X9)|~v3_ordinal1(X8))&(~r1_ordinal1(X8,X9)|r2_hidden(X8,k1_ordinal1(X9))|~v3_ordinal1(X9)|~v3_ordinal1(X8))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t34_ordinal1])])])])).
fof(c_0_21, plain, ![X10]:k1_ordinal1(X10)=k2_xboole_0(X10,k1_tarski(X10)), inference(variable_rename,[status(thm)],[d1_ordinal1])).
fof(c_0_22, plain, ![X11, X12]:(~v3_ordinal1(X11)|~v3_ordinal1(X12)|(r1_ordinal1(X11,X12)|r1_ordinal1(X12,X11))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[connectedness_r1_ordinal1])])).
fof(c_0_23, negated_conjecture, ![X7]:(v3_ordinal1(esk1_0)&(((v3_ordinal1(esk2_0)|~v4_ordinal1(esk1_0))&((r2_hidden(esk2_0,esk1_0)|~v4_ordinal1(esk1_0))&(~r2_hidden(k1_ordinal1(esk2_0),esk1_0)|~v4_ordinal1(esk1_0))))&(v4_ordinal1(esk1_0)|(~v3_ordinal1(X7)|(~r2_hidden(X7,esk1_0)|r2_hidden(k1_ordinal1(X7),esk1_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])])).
fof(c_0_24, plain, ![X13]:(v3_ordinal1(esk3_1(X13))&~r2_hidden(esk3_1(X13),X13)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t38_ordinal1])])])).
cnf(c_0_25, plain, (r2_hidden(X1,k1_ordinal1(X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (r1_ordinal1(X1,X2)|r1_ordinal1(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, negated_conjecture, (v3_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (~r2_hidden(esk3_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, plain, (r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2)), inference(rw,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_31, plain, (v3_ordinal1(esk3_1(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, negated_conjecture, (r1_ordinal1(X1,esk1_0)|r1_ordinal1(esk1_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
fof(c_0_33, plain, ![X15, X16]:((~r1_ordinal1(X15,X16)|r1_tarski(X15,X16)|(~v3_ordinal1(X15)|~v3_ordinal1(X16)))&(~r1_tarski(X15,X16)|r1_ordinal1(X15,X16)|(~v3_ordinal1(X15)|~v3_ordinal1(X16)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r1_ordinal1])])])).
cnf(c_0_34, plain, (~r1_ordinal1(esk3_1(k2_xboole_0(X1,k1_tarski(X1))),X1)|~v3_ordinal1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])])).
cnf(c_0_35, negated_conjecture, (r1_ordinal1(esk1_0,esk3_1(X1))|r1_ordinal1(esk3_1(X1),esk1_0)), inference(spm,[status(thm)],[c_0_32, c_0_31])).
fof(c_0_36, plain, ![X17, X18, X19]:((~v1_ordinal1(X17)|(~r2_hidden(X18,X17)|r1_tarski(X18,X17)))&((r2_hidden(esk4_1(X19),X19)|v1_ordinal1(X19))&(~r1_tarski(esk4_1(X19),X19)|v1_ordinal1(X19)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_ordinal1])])])])])])).
fof(c_0_37, plain, ![X21]:((v1_ordinal1(X21)|~v3_ordinal1(X21))&(v2_ordinal1(X21)|~v3_ordinal1(X21))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_ordinal1])])])).
fof(c_0_38, plain, ![X22, X23, X24, X25, X26]:((~r1_tarski(X22,X23)|(~r2_hidden(X24,X22)|r2_hidden(X24,X23)))&((r2_hidden(esk5_2(X25,X26),X25)|r1_tarski(X25,X26))&(~r2_hidden(esk5_2(X25,X26),X26)|r1_tarski(X25,X26)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_39, plain, (r1_tarski(X1,X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_40, negated_conjecture, (r1_ordinal1(esk1_0,esk3_1(k2_xboole_0(esk1_0,k1_tarski(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_28])])).
fof(c_0_41, plain, ![X28, X29]:(~r2_hidden(X28,X29)|~r1_tarski(X29,X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_42, plain, (r1_tarski(X2,X1)|~v1_ordinal1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_43, plain, (v1_ordinal1(X1)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_44, plain, ![X30, X31, X32, X34, X35, X36, X37, X39]:((((r2_hidden(X32,esk6_3(X30,X31,X32))|~r2_hidden(X32,X31)|X31!=k3_tarski(X30))&(r2_hidden(esk6_3(X30,X31,X32),X30)|~r2_hidden(X32,X31)|X31!=k3_tarski(X30)))&(~r2_hidden(X34,X35)|~r2_hidden(X35,X30)|r2_hidden(X34,X31)|X31!=k3_tarski(X30)))&((~r2_hidden(esk7_2(X36,X37),X37)|(~r2_hidden(esk7_2(X36,X37),X39)|~r2_hidden(X39,X36))|X37=k3_tarski(X36))&((r2_hidden(esk7_2(X36,X37),esk8_2(X36,X37))|r2_hidden(esk7_2(X36,X37),X37)|X37=k3_tarski(X36))&(r2_hidden(esk8_2(X36,X37),X36)|r2_hidden(esk7_2(X36,X37),X37)|X37=k3_tarski(X36))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_tarski])])])])])])).
fof(c_0_45, plain, ![X41, X42]:(~v3_ordinal1(X42)|(~r2_hidden(X41,X42)|v3_ordinal1(X41))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_ordinal1])])).
cnf(c_0_46, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_47, negated_conjecture, (r1_tarski(esk1_0,esk3_1(k2_xboole_0(esk1_0,k1_tarski(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_31]), c_0_28])])).
cnf(c_0_48, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_49, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,X2)|~v3_ordinal1(X2)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_50, plain, (r2_hidden(esk6_3(X1,X2,X3),X1)|~r2_hidden(X3,X2)|X2!=k3_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_51, plain, ![X43]:((r2_hidden(esk9_1(X43),X43)|v3_ordinal1(k3_tarski(X43)))&(~v3_ordinal1(esk9_1(X43))|v3_ordinal1(k3_tarski(X43)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t35_ordinal1])])])])).
cnf(c_0_52, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_53, negated_conjecture, (r2_hidden(X1,esk3_1(k2_xboole_0(esk1_0,k1_tarski(esk1_0))))|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_54, plain, (~r2_hidden(X1,X2)|~r2_hidden(X2,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_55, plain, (r2_hidden(esk6_3(X1,k3_tarski(X1),X2),X1)|~r2_hidden(X2,k3_tarski(X1))), inference(er,[status(thm)],[c_0_50])).
cnf(c_0_56, plain, (r2_hidden(X1,esk6_3(X2,X3,X1))|~r2_hidden(X1,X3)|X3!=k3_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_57, plain, ![X45, X46]:(~v3_ordinal1(X45)|(~v3_ordinal1(X46)|(r2_hidden(X45,X46)|(X45=X46|r2_hidden(X46,X45))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_ordinal1])])])])).
cnf(c_0_58, plain, (v3_ordinal1(k3_tarski(X1))|~v3_ordinal1(esk9_1(X1))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_59, negated_conjecture, (v3_ordinal1(X1)|~r2_hidden(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_31])])).
cnf(c_0_60, plain, (~r2_hidden(X1,esk6_3(X1,k3_tarski(X1),X2))|~r2_hidden(X2,k3_tarski(X1))|~v3_ordinal1(esk6_3(X1,k3_tarski(X1),X2))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_61, plain, (r2_hidden(X1,esk6_3(X2,k3_tarski(X2),X1))|~r2_hidden(X1,k3_tarski(X2))), inference(er,[status(thm)],[c_0_56])).
cnf(c_0_62, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X2,X3)|X4!=k3_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_63, negated_conjecture, (v4_ordinal1(esk1_0)|r2_hidden(k1_ordinal1(X1),esk1_0)|~v3_ordinal1(X1)|~r2_hidden(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_64, plain, ![X47]:r2_hidden(X47,k1_ordinal1(X47)), inference(variable_rename,[status(thm)],[t10_ordinal1])).
cnf(c_0_65, plain, (r2_hidden(X1,X2)|X1=X2|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_66, negated_conjecture, (v3_ordinal1(k3_tarski(X1))|~r2_hidden(esk9_1(X1),esk1_0)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_67, plain, (r2_hidden(esk9_1(X1),X1)|v3_ordinal1(k3_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_68, plain, (~r2_hidden(X1,k3_tarski(X1))|~v3_ordinal1(esk6_3(X1,k3_tarski(X1),X1))), inference(spm,[status(thm)],[c_0_60, c_0_61])).
fof(c_0_69, plain, ![X48, X49]:(((r1_tarski(X48,X49)|X48!=X49)&(r1_tarski(X49,X48)|X48!=X49))&(~r1_tarski(X48,X49)|~r1_tarski(X49,X48)|X48=X49)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_70, plain, ![X50, X51]:((~r2_hidden(X50,X51)|r1_ordinal1(k1_ordinal1(X50),X51)|~v3_ordinal1(X51)|~v3_ordinal1(X50))&(~r1_ordinal1(k1_ordinal1(X50),X51)|r2_hidden(X50,X51)|~v3_ordinal1(X51)|~v3_ordinal1(X50))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t33_ordinal1])])])])).
cnf(c_0_71, plain, (r2_hidden(X1,k3_tarski(X2))|~r2_hidden(X3,X2)|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_62])).
cnf(c_0_72, negated_conjecture, (v4_ordinal1(esk1_0)|r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),esk1_0)|~v3_ordinal1(X1)|~r2_hidden(X1,esk1_0)), inference(rw,[status(thm)],[c_0_63, c_0_26])).
cnf(c_0_73, plain, (r2_hidden(X1,k1_ordinal1(X1))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_74, negated_conjecture, (X1=esk1_0|r2_hidden(X1,esk1_0)|r2_hidden(esk1_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_65, c_0_28])).
cnf(c_0_75, negated_conjecture, (v3_ordinal1(k3_tarski(esk1_0))), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_76, negated_conjecture, (~r2_hidden(esk6_3(X1,k3_tarski(X1),X1),esk1_0)|~r2_hidden(X1,k3_tarski(X1))), inference(spm,[status(thm)],[c_0_68, c_0_59])).
cnf(c_0_77, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_78, plain, (r1_ordinal1(k1_ordinal1(X1),X2)|~r2_hidden(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
fof(c_0_79, plain, ![X52]:(~v3_ordinal1(X52)|v3_ordinal1(k1_ordinal1(X52))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t29_ordinal1])])).
fof(c_0_80, plain, ![X53]:((~v4_ordinal1(X53)|X53=k3_tarski(X53))&(X53!=k3_tarski(X53)|v4_ordinal1(X53))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_ordinal1])])).
cnf(c_0_81, negated_conjecture, (r2_hidden(X1,k3_tarski(esk1_0))|v4_ordinal1(esk1_0)|~r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(X2,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_59])).
cnf(c_0_82, plain, (r2_hidden(X1,k2_xboole_0(X1,k1_tarski(X1)))), inference(rw,[status(thm)],[c_0_73, c_0_26])).
cnf(c_0_83, negated_conjecture, (k3_tarski(esk1_0)=esk1_0|r2_hidden(esk1_0,k3_tarski(esk1_0))|r2_hidden(k3_tarski(esk1_0),esk1_0)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_84, negated_conjecture, (~r2_hidden(esk1_0,k3_tarski(esk1_0))), inference(spm,[status(thm)],[c_0_76, c_0_55])).
cnf(c_0_85, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_77])).
cnf(c_0_86, plain, (r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_78, c_0_26])).
cnf(c_0_87, plain, (v3_ordinal1(k1_ordinal1(X1))|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_88, negated_conjecture, (v3_ordinal1(esk2_0)|~v4_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_89, plain, (v4_ordinal1(X1)|X1!=k3_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_90, negated_conjecture, (r2_hidden(X1,k3_tarski(esk1_0))|v4_ordinal1(esk1_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_91, negated_conjecture, (k3_tarski(esk1_0)=esk1_0|r2_hidden(k3_tarski(esk1_0),esk1_0)), inference(sr,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_92, plain, (~r2_hidden(X1,X1)), inference(spm,[status(thm)],[c_0_48, c_0_85])).
cnf(c_0_93, plain, (X1=k3_tarski(X1)|~v4_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_94, plain, (r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X2)|~r2_hidden(X1,X2)|~v3_ordinal1(X2)), inference(csr,[status(thm)],[c_0_86, c_0_52])).
cnf(c_0_95, negated_conjecture, (~r2_hidden(k1_ordinal1(esk2_0),esk1_0)|~v4_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_96, plain, (v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)), inference(rw,[status(thm)],[c_0_87, c_0_26])).
cnf(c_0_97, negated_conjecture, (v3_ordinal1(esk2_0)|k3_tarski(esk1_0)!=esk1_0), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_98, negated_conjecture, (k3_tarski(esk1_0)=esk1_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_92]), c_0_93])).
cnf(c_0_99, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v4_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_100, plain, (r1_tarski(k2_xboole_0(X1,k1_tarski(X1)),X2)|~r2_hidden(X1,X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X2)), inference(spm,[status(thm)],[c_0_39, c_0_94])).
cnf(c_0_101, negated_conjecture, (~v4_ordinal1(esk1_0)|~r2_hidden(k2_xboole_0(esk2_0,k1_tarski(esk2_0)),esk1_0)), inference(rw,[status(thm)],[c_0_95, c_0_26])).
cnf(c_0_102, negated_conjecture, (k2_xboole_0(X1,k1_tarski(X1))=esk1_0|r2_hidden(esk1_0,k2_xboole_0(X1,k1_tarski(X1)))|r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),esk1_0)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_74, c_0_96])).
cnf(c_0_103, negated_conjecture, (v3_ordinal1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_98])])).
cnf(c_0_104, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|k3_tarski(esk1_0)!=esk1_0), inference(spm,[status(thm)],[c_0_99, c_0_89])).
cnf(c_0_105, plain, (~r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(X2,X1)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_100]), c_0_52])).
cnf(c_0_106, negated_conjecture, (k2_xboole_0(esk2_0,k1_tarski(esk2_0))=esk1_0|r2_hidden(esk1_0,k2_xboole_0(esk2_0,k1_tarski(esk2_0)))|~v4_ordinal1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_103])])).
cnf(c_0_107, negated_conjecture, (r2_hidden(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_104, c_0_98])])).
cnf(c_0_108, negated_conjecture, (k2_xboole_0(esk2_0,k1_tarski(esk2_0))=esk1_0|~v4_ordinal1(esk1_0)|~v3_ordinal1(k2_xboole_0(esk2_0,k1_tarski(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_107])])).
cnf(c_0_109, negated_conjecture, (k2_xboole_0(esk2_0,k1_tarski(esk2_0))=esk1_0|~v4_ordinal1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_96]), c_0_103])])).
cnf(c_0_110, negated_conjecture, (~r2_hidden(X1,esk1_0)|~r2_hidden(esk2_0,X1)|~v4_ordinal1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_109]), c_0_28])])).
cnf(c_0_111, negated_conjecture, (~r2_hidden(X1,esk1_0)|~r2_hidden(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_89]), c_0_98])])).
cnf(c_0_112, negated_conjecture, (~r2_hidden(esk2_0,esk6_3(esk1_0,esk1_0,X1))|~r2_hidden(X1,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_55]), c_0_98]), c_0_98])).
cnf(c_0_113, negated_conjecture, (r2_hidden(X1,esk6_3(esk1_0,esk1_0,X1))|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_61, c_0_98])).
cnf(c_0_114, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_107])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 115
# Proof object clause steps            : 76
# Proof object formula steps           : 39
# Proof object conjectures             : 39
# Proof object clause conjectures      : 36
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 34
# Proof object simplifying inferences  : 43
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 41
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 40
# Processed clauses                    : 866
# ...of these trivial                  : 1
# ...subsumed                          : 427
# ...remaining for further processing  : 438
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 28
# Backward-rewritten                   : 31
# Generated clauses                    : 2370
# ...of the previous two non-trivial   : 2218
# Contextual simplify-reflections      : 51
# Paramodulations                      : 2363
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 333
#    Positive orientable unit clauses  : 18
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 20
#    Non-unit-clauses                  : 295
# Current number of unprocessed clauses: 1391
# ...number of literals in the above   : 5778
# Current number of archived formulas  : 0
# Current number of archived clauses   : 101
# Clause-clause subsumption calls (NU) : 11808
# Rec. Clause-clause subsumption calls : 5874
# Non-unit clause-clause subsumptions  : 220
# Unit Clause-clause subsumption calls : 526
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 39899

# -------------------------------------------------
# User time                : 0.030 s
# System time              : 0.010 s
# Total time               : 0.040 s
# Maximum resident set size: 28736 pages
