# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/0c645edc5bbe49679ab248aa641a97df.E_input', d10_xboole_0)).
fof(t21_xboole_1, axiom, ![X1, X2]:k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/0c645edc5bbe49679ab248aa641a97df.E_input', t21_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/0c645edc5bbe49679ab248aa641a97df.E_input', t48_xboole_1)).
fof(t106_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k4_xboole_0(X2,X3))=>(r1_tarski(X1,X2)&r1_xboole_0(X1,X3))), file('data/MPTP/proofs/0c645edc5bbe49679ab248aa641a97df.E_input', t106_xboole_1)).
fof(l1_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/0c645edc5bbe49679ab248aa641a97df.E_input', l1_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/0c645edc5bbe49679ab248aa641a97df.E_input', t69_enumset1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/0c645edc5bbe49679ab248aa641a97df.E_input', commutativity_k3_xboole_0)).
fof(t40_xboole_1, axiom, ![X1, X2]:k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/0c645edc5bbe49679ab248aa641a97df.E_input', t40_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/0c645edc5bbe49679ab248aa641a97df.E_input', commutativity_k2_xboole_0)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/0c645edc5bbe49679ab248aa641a97df.E_input', t39_xboole_1)).
fof(t38_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,k4_xboole_0(X2,X1))=>X1=k1_xboole_0), file('data/MPTP/proofs/0c645edc5bbe49679ab248aa641a97df.E_input', t38_xboole_1)).
fof(t63_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_xboole_0(X2,X3))=>r1_xboole_0(X1,X3)), file('data/MPTP/proofs/0c645edc5bbe49679ab248aa641a97df.E_input', t63_xboole_1)).
fof(t103_zfmisc_1, conjecture, ![X1, X2, X3, X4]:~((r1_tarski(X1,k2_zfmisc_1(X2,X3))&(r2_hidden(X4,X1)&![X5, X6]:~((r2_hidden(X5,X2)&(r2_hidden(X6,X3)&X4=k4_tarski(X5,X6))))))), file('data/MPTP/proofs/0c645edc5bbe49679ab248aa641a97df.E_input', t103_zfmisc_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/0c645edc5bbe49679ab248aa641a97df.E_input', t3_boole)).
fof(t55_zfmisc_1, axiom, ![X1, X2, X3]:~((r1_xboole_0(k2_tarski(X1,X2),X3)&r2_hidden(X1,X3))), file('data/MPTP/proofs/0c645edc5bbe49679ab248aa641a97df.E_input', t55_zfmisc_1)).
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/0c645edc5bbe49679ab248aa641a97df.E_input', d2_zfmisc_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/0c645edc5bbe49679ab248aa641a97df.E_input', d5_xboole_0)).
fof(c_0_17, plain, ![X14, X15]:(((r1_tarski(X14,X15)|X14!=X15)&(r1_tarski(X15,X14)|X14!=X15))&(~r1_tarski(X14,X15)|~r1_tarski(X15,X14)|X14=X15)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_18, plain, ![X7, X8]:k3_xboole_0(X7,k2_xboole_0(X7,X8))=X7, inference(variable_rename,[status(thm)],[t21_xboole_1])).
fof(c_0_19, plain, ![X9, X10]:k4_xboole_0(X9,k4_xboole_0(X9,X10))=k3_xboole_0(X9,X10), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_20, plain, ![X16, X17, X18]:((r1_tarski(X16,X17)|~r1_tarski(X16,k4_xboole_0(X17,X18)))&(r1_xboole_0(X16,X18)|~r1_tarski(X16,k4_xboole_0(X17,X18)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t106_xboole_1])])])).
cnf(c_0_21, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_22, plain, ![X33, X34]:((~r1_tarski(k1_tarski(X33),X34)|r2_hidden(X33,X34))&(~r2_hidden(X33,X34)|r1_tarski(k1_tarski(X33),X34))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])).
fof(c_0_23, plain, ![X35]:k2_tarski(X35,X35)=k1_tarski(X35), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_24, plain, ![X36, X37]:k3_xboole_0(X36,X37)=k3_xboole_0(X37,X36), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_25, plain, ![X12, X13]:k4_xboole_0(k2_xboole_0(X12,X13),X13)=k4_xboole_0(X12,X13), inference(variable_rename,[status(thm)],[t40_xboole_1])).
fof(c_0_26, plain, ![X38, X39]:k2_xboole_0(X38,X39)=k2_xboole_0(X39,X38), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_27, plain, ![X40, X41]:k2_xboole_0(X40,k4_xboole_0(X41,X40))=k2_xboole_0(X40,X41), inference(variable_rename,[status(thm)],[t39_xboole_1])).
fof(c_0_28, plain, ![X19, X20]:(~r1_tarski(X19,k4_xboole_0(X20,X19))|X19=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t38_xboole_1])])).
cnf(c_0_29, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_30, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_31, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k4_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_32, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_21])).
fof(c_0_33, plain, ![X42, X43, X44]:(~r1_tarski(X42,X43)|~r1_xboole_0(X43,X44)|r1_xboole_0(X42,X44)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_xboole_1])])).
cnf(c_0_34, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_35, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_36, negated_conjecture, ~(![X1, X2, X3, X4]:~((r1_tarski(X1,k2_zfmisc_1(X2,X3))&(r2_hidden(X4,X1)&![X5, X6]:~((r2_hidden(X5,X2)&(r2_hidden(X6,X3)&X4=k4_tarski(X5,X6)))))))), inference(assume_negation,[status(cth)],[t103_zfmisc_1])).
cnf(c_0_37, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_38, plain, (k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_39, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_40, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_41, plain, (X1=k1_xboole_0|~r1_tarski(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_42, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k2_xboole_0(X1,X2)))=X1), inference(rw,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_43, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
fof(c_0_44, plain, ![X11]:k4_xboole_0(X11,k1_xboole_0)=X11, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_45, plain, (r1_xboole_0(X1,X3)|~r1_tarski(X1,X2)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_46, plain, (r1_tarski(k2_tarski(X1,X1),X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_34, c_0_35])).
fof(c_0_47, negated_conjecture, ![X49, X50]:(r1_tarski(esk2_0,k2_zfmisc_1(esk3_0,esk4_0))&(r2_hidden(esk5_0,esk2_0)&(~r2_hidden(X49,esk3_0)|(~r2_hidden(X50,esk4_0)|esk5_0!=k4_tarski(X49,X50))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_36])])])])).
cnf(c_0_48, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k4_xboole_0(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_49, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_30]), c_0_30])).
cnf(c_0_50, plain, (k4_xboole_0(k2_xboole_0(X1,X2),X1)=k4_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_51, plain, (k4_xboole_0(k2_xboole_0(X1,X2),k4_xboole_0(X2,X1))=k4_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_38, c_0_40])).
cnf(c_0_52, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43])])).
cnf(c_0_53, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_54, plain, ![X30, X31, X32]:(~r1_xboole_0(k2_tarski(X30,X31),X32)|~r2_hidden(X30,X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_zfmisc_1])])).
cnf(c_0_55, plain, (r1_xboole_0(k2_tarski(X1,X1),X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_56, negated_conjecture, (r2_hidden(esk5_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_57, negated_conjecture, (r1_tarski(esk2_0,k2_zfmisc_1(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_58, plain, (r1_xboole_0(k4_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_48, c_0_32])).
cnf(c_0_59, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51]), c_0_52]), c_0_53])).
fof(c_0_60, plain, ![X51, X52, X53, X54, X57, X58, X59, X60, X61, X62, X64, X65]:((((r2_hidden(esk6_4(X51,X52,X53,X54),X51)|~r2_hidden(X54,X53)|X53!=k2_zfmisc_1(X51,X52))&((r2_hidden(esk7_4(X51,X52,X53,X54),X52)|~r2_hidden(X54,X53)|X53!=k2_zfmisc_1(X51,X52))&(X54=k4_tarski(esk6_4(X51,X52,X53,X54),esk7_4(X51,X52,X53,X54))|~r2_hidden(X54,X53)|X53!=k2_zfmisc_1(X51,X52))))&(~r2_hidden(X58,X51)|(~r2_hidden(X59,X52)|X57!=k4_tarski(X58,X59))|r2_hidden(X57,X53)|X53!=k2_zfmisc_1(X51,X52)))&((~r2_hidden(esk8_3(X60,X61,X62),X62)|(~r2_hidden(X64,X60)|(~r2_hidden(X65,X61)|esk8_3(X60,X61,X62)!=k4_tarski(X64,X65)))|X62=k2_zfmisc_1(X60,X61))&((r2_hidden(esk9_3(X60,X61,X62),X60)|r2_hidden(esk8_3(X60,X61,X62),X62)|X62=k2_zfmisc_1(X60,X61))&((r2_hidden(esk10_3(X60,X61,X62),X61)|r2_hidden(esk8_3(X60,X61,X62),X62)|X62=k2_zfmisc_1(X60,X61))&(esk8_3(X60,X61,X62)=k4_tarski(esk9_3(X60,X61,X62),esk10_3(X60,X61,X62))|r2_hidden(esk8_3(X60,X61,X62),X62)|X62=k2_zfmisc_1(X60,X61)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
cnf(c_0_61, plain, (~r1_xboole_0(k2_tarski(X1,X2),X3)|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_62, negated_conjecture, (r1_xboole_0(k2_tarski(esk5_0,esk5_0),X1)|~r1_xboole_0(esk2_0,X1)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_63, negated_conjecture, (r1_xboole_0(esk2_0,X1)|~r1_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),X1)), inference(spm,[status(thm)],[c_0_45, c_0_57])).
cnf(c_0_64, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_58, c_0_59])).
fof(c_0_65, plain, ![X21, X22, X23, X24, X25, X26, X27, X28]:((((r2_hidden(X24,X21)|~r2_hidden(X24,X23)|X23!=k4_xboole_0(X21,X22))&(~r2_hidden(X24,X22)|~r2_hidden(X24,X23)|X23!=k4_xboole_0(X21,X22)))&(~r2_hidden(X25,X21)|r2_hidden(X25,X22)|r2_hidden(X25,X23)|X23!=k4_xboole_0(X21,X22)))&((~r2_hidden(esk1_3(X26,X27,X28),X28)|(~r2_hidden(esk1_3(X26,X27,X28),X26)|r2_hidden(esk1_3(X26,X27,X28),X27))|X28=k4_xboole_0(X26,X27))&((r2_hidden(esk1_3(X26,X27,X28),X26)|r2_hidden(esk1_3(X26,X27,X28),X28)|X28=k4_xboole_0(X26,X27))&(~r2_hidden(esk1_3(X26,X27,X28),X27)|r2_hidden(esk1_3(X26,X27,X28),X28)|X28=k4_xboole_0(X26,X27))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_66, plain, (X1=k4_tarski(esk6_4(X2,X3,X4,X1),esk7_4(X2,X3,X4,X1))|~r2_hidden(X1,X4)|X4!=k2_zfmisc_1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_67, negated_conjecture, (~r2_hidden(esk5_0,X1)|~r1_xboole_0(esk2_0,X1)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_68, negated_conjecture, (r1_xboole_0(esk2_0,k4_xboole_0(X1,k2_zfmisc_1(esk3_0,esk4_0)))), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_69, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_70, plain, (r2_hidden(X1,X2)|~r1_tarski(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_71, negated_conjecture, (~r2_hidden(X1,esk3_0)|~r2_hidden(X2,esk4_0)|esk5_0!=k4_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_72, plain, (k4_tarski(esk6_4(X1,X2,k2_zfmisc_1(X1,X2),X3),esk7_4(X1,X2,k2_zfmisc_1(X1,X2),X3))=X3|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_66])).
cnf(c_0_73, plain, (r2_hidden(esk7_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_74, negated_conjecture, (~r2_hidden(esk5_0,k4_xboole_0(X1,k2_zfmisc_1(esk3_0,esk4_0)))), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_75, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_69])).
cnf(c_0_76, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_tarski(X1,X1),X2)), inference(rw,[status(thm)],[c_0_70, c_0_35])).
cnf(c_0_77, negated_conjecture, (~r2_hidden(esk7_4(X1,X2,k2_zfmisc_1(X1,X2),esk5_0),esk4_0)|~r2_hidden(esk6_4(X1,X2,k2_zfmisc_1(X1,X2),esk5_0),esk3_0)|~r2_hidden(esk5_0,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72])])).
cnf(c_0_78, plain, (r2_hidden(esk7_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X2)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_73])).
cnf(c_0_79, plain, (r2_hidden(esk6_4(X1,X2,X3,X4),X1)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_80, negated_conjecture, (r2_hidden(esk5_0,k2_zfmisc_1(esk3_0,esk4_0))|~r2_hidden(esk5_0,X1)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_81, plain, (r2_hidden(X1,k2_tarski(X1,X1))), inference(spm,[status(thm)],[c_0_76, c_0_32])).
cnf(c_0_82, negated_conjecture, (~r2_hidden(esk6_4(X1,esk4_0,k2_zfmisc_1(X1,esk4_0),esk5_0),esk3_0)|~r2_hidden(esk5_0,k2_zfmisc_1(X1,esk4_0))), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_83, plain, (r2_hidden(esk6_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X1)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_79])).
cnf(c_0_84, negated_conjecture, (r2_hidden(esk5_0,k2_zfmisc_1(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_85, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_84])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 86
# Proof object clause steps            : 51
# Proof object formula steps           : 35
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 33
# Processed clauses                    : 1102
# ...of these trivial                  : 46
# ...subsumed                          : 779
# ...remaining for further processing  : 277
# Other redundant clauses eliminated   : 11
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 63
# Generated clauses                    : 3782
# ...of the previous two non-trivial   : 3167
# Contextual simplify-reflections      : 1
# Paramodulations                      : 3766
# Factorizations                       : 6
# Equation resolutions                 : 11
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 170
#    Positive orientable unit clauses  : 51
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 8
#    Non-unit-clauses                  : 109
# Current number of unprocessed clauses: 1878
# ...number of literals in the above   : 4467
# Current number of archived formulas  : 0
# Current number of archived clauses   : 100
# Clause-clause subsumption calls (NU) : 3072
# Rec. Clause-clause subsumption calls : 2202
# Non-unit clause-clause subsumptions  : 266
# Unit Clause-clause subsumption calls : 357
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 105
# BW rewrite match successes           : 42
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 37076

# -------------------------------------------------
# User time                : 0.030 s
# System time              : 0.003 s
# Total time               : 0.033 s
# Maximum resident set size: 28632 pages
