# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/0c862a57a4644455bd83564444828c34.E_input', d2_zfmisc_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/0c862a57a4644455bd83564444828c34.E_input', t113_zfmisc_1)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/0c862a57a4644455bd83564444828c34.E_input', t7_mcart_1)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/0c862a57a4644455bd83564444828c34.E_input', d2_tarski)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/0c862a57a4644455bd83564444828c34.E_input', t70_enumset1)).
fof(t51_mcart_1, conjecture, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(~(X4=k5_mcart_1(X1,X2,X3,X4))&(~(X4=k6_mcart_1(X1,X2,X3,X4))&~(X4=k7_mcart_1(X1,X2,X3,X4)))))))))), file('data/MPTP/proofs/0c862a57a4644455bd83564444828c34.E_input', t51_mcart_1)).
fof(t48_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>X4=k3_mcart_1(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4),k7_mcart_1(X1,X2,X3,X4)))))))), file('data/MPTP/proofs/0c862a57a4644455bd83564444828c34.E_input', t48_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/0c862a57a4644455bd83564444828c34.E_input', d3_zfmisc_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/0c862a57a4644455bd83564444828c34.E_input', d3_mcart_1)).
fof(t20_mcart_1, axiom, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1)))), file('data/MPTP/proofs/0c862a57a4644455bd83564444828c34.E_input', t20_mcart_1)).
fof(t29_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4, X5]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k3_mcart_1(X3,X4,X5))))))), file('data/MPTP/proofs/0c862a57a4644455bd83564444828c34.E_input', t29_mcart_1)).
fof(c_0_11, plain, ![X20, X21, X22, X23, X26, X27, X28, X29, X30, X31, X33, X34]:((((r2_hidden(esk6_4(X20,X21,X22,X23),X20)|~r2_hidden(X23,X22)|X22!=k2_zfmisc_1(X20,X21))&((r2_hidden(esk7_4(X20,X21,X22,X23),X21)|~r2_hidden(X23,X22)|X22!=k2_zfmisc_1(X20,X21))&(X23=k4_tarski(esk6_4(X20,X21,X22,X23),esk7_4(X20,X21,X22,X23))|~r2_hidden(X23,X22)|X22!=k2_zfmisc_1(X20,X21))))&(~r2_hidden(X27,X20)|(~r2_hidden(X28,X21)|X26!=k4_tarski(X27,X28))|r2_hidden(X26,X22)|X22!=k2_zfmisc_1(X20,X21)))&((~r2_hidden(esk8_3(X29,X30,X31),X31)|(~r2_hidden(X33,X29)|(~r2_hidden(X34,X30)|esk8_3(X29,X30,X31)!=k4_tarski(X33,X34)))|X31=k2_zfmisc_1(X29,X30))&((r2_hidden(esk9_3(X29,X30,X31),X29)|r2_hidden(esk8_3(X29,X30,X31),X31)|X31=k2_zfmisc_1(X29,X30))&((r2_hidden(esk10_3(X29,X30,X31),X30)|r2_hidden(esk8_3(X29,X30,X31),X31)|X31=k2_zfmisc_1(X29,X30))&(esk8_3(X29,X30,X31)=k4_tarski(esk9_3(X29,X30,X31),esk10_3(X29,X30,X31))|r2_hidden(esk8_3(X29,X30,X31),X31)|X31=k2_zfmisc_1(X29,X30)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
fof(c_0_12, plain, ![X39, X40]:((k2_zfmisc_1(X39,X40)!=k1_xboole_0|(X39=k1_xboole_0|X40=k1_xboole_0))&((X39!=k1_xboole_0|k2_zfmisc_1(X39,X40)=k1_xboole_0)&(X40!=k1_xboole_0|k2_zfmisc_1(X39,X40)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_13, plain, (X1=k4_tarski(esk6_4(X2,X3,X4,X1),esk7_4(X2,X3,X4,X1))|~r2_hidden(X1,X4)|X4!=k2_zfmisc_1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_14, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15, plain, (r2_hidden(esk6_4(X1,X2,X3,X4),X1)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_17, plain, ![X37, X38]:(k1_mcart_1(k4_tarski(X37,X38))=X37&k2_mcart_1(k4_tarski(X37,X38))=X38), inference(variable_rename,[status(thm)],[t7_mcart_1])).
cnf(c_0_18, plain, (k4_tarski(esk6_4(X1,X2,k2_zfmisc_1(X1,X2),X3),esk7_4(X1,X2,k2_zfmisc_1(X1,X2),X3))=X3|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_14])).
fof(c_0_20, plain, ![X11, X12, X13, X14, X15, X16, X17, X18]:(((~r2_hidden(X14,X13)|(X14=X11|X14=X12)|X13!=k2_tarski(X11,X12))&((X15!=X11|r2_hidden(X15,X13)|X13!=k2_tarski(X11,X12))&(X15!=X12|r2_hidden(X15,X13)|X13!=k2_tarski(X11,X12))))&(((esk5_3(X16,X17,X18)!=X16|~r2_hidden(esk5_3(X16,X17,X18),X18)|X18=k2_tarski(X16,X17))&(esk5_3(X16,X17,X18)!=X17|~r2_hidden(esk5_3(X16,X17,X18),X18)|X18=k2_tarski(X16,X17)))&(r2_hidden(esk5_3(X16,X17,X18),X18)|(esk5_3(X16,X17,X18)=X16|esk5_3(X16,X17,X18)=X17)|X18=k2_tarski(X16,X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_21, plain, ![X41, X42]:k1_enumset1(X41,X41,X42)=k2_tarski(X41,X42), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_22, plain, (r2_hidden(esk6_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X1)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (k4_tarski(esk6_4(k1_xboole_0,X1,k1_xboole_0,X2),esk7_4(k1_xboole_0,X1,k1_xboole_0,X2))=X2|~r2_hidden(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_26, plain, (r2_hidden(X5,X6)|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)|X5!=k4_tarski(X1,X3)|X6!=k2_zfmisc_1(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_27, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, plain, (r2_hidden(esk6_4(X1,k1_xboole_0,k1_xboole_0,X2),X1)|~r2_hidden(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_30, plain, (esk6_4(k1_xboole_0,X1,k1_xboole_0,X2)=k1_mcart_1(X2)|~r2_hidden(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31, plain, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_26])])).
cnf(c_0_32, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X4,X2)), inference(rw,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_33, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_34, plain, (r2_hidden(esk7_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_35, plain, (r2_hidden(k1_mcart_1(X1),k1_xboole_0)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_36, plain, (r2_hidden(k4_tarski(X1,X2),k1_xboole_0)|~r2_hidden(X2,k1_xboole_0)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_31, c_0_23])).
cnf(c_0_37, plain, (r2_hidden(X1,k1_enumset1(X2,X2,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_32])])).
cnf(c_0_38, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X2,X2,X4)), inference(rw,[status(thm)],[c_0_33, c_0_28])).
cnf(c_0_39, plain, (r2_hidden(esk5_3(X1,X2,X3),X3)|esk5_3(X1,X2,X3)=X1|esk5_3(X1,X2,X3)=X2|X3=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_40, plain, (r2_hidden(esk7_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X2)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_34])).
cnf(c_0_41, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_42, plain, (r2_hidden(X1,k1_xboole_0)|~r2_hidden(k4_tarski(X1,X2),k1_xboole_0)), inference(spm,[status(thm)],[c_0_35, c_0_24])).
cnf(c_0_43, plain, (r2_hidden(k4_tarski(X1,X2),k1_xboole_0)|~r2_hidden(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_44, plain, (r2_hidden(X1,k1_enumset1(X1,X1,X2))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_38])])).
cnf(c_0_45, plain, (X3=k1_enumset1(X1,X1,X2)|esk5_3(X1,X2,X3)=X2|esk5_3(X1,X2,X3)=X1|r2_hidden(esk5_3(X1,X2,X3),X3)), inference(rw,[status(thm)],[c_0_39, c_0_28])).
fof(c_0_46, negated_conjecture, ~(![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(~(X4=k5_mcart_1(X1,X2,X3,X4))&(~(X4=k6_mcart_1(X1,X2,X3,X4))&~(X4=k7_mcart_1(X1,X2,X3,X4))))))))))), inference(assume_negation,[status(cth)],[t51_mcart_1])).
fof(c_0_47, plain, ![X57, X58, X59, X60]:(X57=k1_xboole_0|(X58=k1_xboole_0|(X59=k1_xboole_0|(~m1_subset_1(X60,k3_zfmisc_1(X57,X58,X59))|X60=k3_mcart_1(k5_mcart_1(X57,X58,X59,X60),k6_mcart_1(X57,X58,X59,X60),k7_mcart_1(X57,X58,X59,X60)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t48_mcart_1])])])])).
fof(c_0_48, plain, ![X51, X52, X53]:k3_zfmisc_1(X51,X52,X53)=k2_zfmisc_1(k2_zfmisc_1(X51,X52),X53), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_49, plain, ![X54, X55, X56]:k3_mcart_1(X54,X55,X56)=k4_tarski(k4_tarski(X54,X55),X56), inference(variable_rename,[status(thm)],[d3_mcart_1])).
cnf(c_0_50, plain, (r2_hidden(esk7_4(k1_xboole_0,X1,k1_xboole_0,X2),X1)|~r2_hidden(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_40, c_0_19])).
cnf(c_0_51, plain, (esk7_4(k1_xboole_0,X1,k1_xboole_0,X2)=k2_mcart_1(X2)|~r2_hidden(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_41, c_0_25])).
cnf(c_0_52, plain, (r2_hidden(X1,k1_xboole_0)|~r2_hidden(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_53, plain, (esk5_3(X1,X2,X3)=X1|esk5_3(X1,X2,X3)=X2|r2_hidden(esk5_3(X1,X2,X3),X3)|r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
fof(c_0_54, negated_conjecture, (esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&(m1_subset_1(esk4_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))&(esk4_0=k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)|(esk4_0=k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)|esk4_0=k7_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_46])])])])).
fof(c_0_55, plain, ![X48, X49, X50]:((X48!=k1_mcart_1(X48)|X48!=k4_tarski(X49,X50))&(X48!=k2_mcart_1(X48)|X48!=k4_tarski(X49,X50))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t20_mcart_1])])])])])).
cnf(c_0_56, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k3_mcart_1(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4),k7_mcart_1(X1,X2,X3,X4))|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_57, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_58, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_59, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_60, plain, (r2_hidden(k2_mcart_1(X1),X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_61, plain, (esk5_3(X1,X2,k1_xboole_0)=X2|esk5_3(X1,X2,k1_xboole_0)=X1|r2_hidden(X3,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_52])).
cnf(c_0_62, negated_conjecture, (m1_subset_1(esk4_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_63, plain, (X1!=k2_mcart_1(X1)|X1!=k4_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_64, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X4=k4_tarski(k4_tarski(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4)),k7_mcart_1(X1,X2,X3,X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_57]), c_0_58])).
cnf(c_0_65, plain, (X1=X4|X1=X3|X2!=k1_enumset1(X3,X3,X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_59, c_0_28])).
cnf(c_0_66, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k1_xboole_0)), inference(spm,[status(thm)],[c_0_60, c_0_41])).
cnf(c_0_67, plain, (esk5_3(X1,X1,k1_xboole_0)=X1|r2_hidden(X2,k1_xboole_0)), inference(er,[status(thm)],[inference(ef,[status(thm)],[c_0_61])])).
fof(c_0_68, plain, ![X43, X45, X46, X47]:((r2_hidden(esk11_1(X43),X43)|X43=k1_xboole_0)&((~r2_hidden(X45,X43)|esk11_1(X43)!=k3_mcart_1(X45,X46,X47)|X43=k1_xboole_0)&(~r2_hidden(X46,X43)|esk11_1(X43)!=k3_mcart_1(X45,X46,X47)|X43=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t29_mcart_1])])])])])])])).
cnf(c_0_69, negated_conjecture, (esk4_0=k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)|esk4_0=k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)|esk4_0=k7_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_70, negated_conjecture, (m1_subset_1(esk4_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(rw,[status(thm)],[c_0_62, c_0_57])).
cnf(c_0_71, plain, (k4_tarski(X1,X2)!=X2), inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_63]), c_0_41])).
cnf(c_0_72, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_73, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_74, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_75, plain, (k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(spm,[status(thm)],[c_0_41, c_0_64])).
cnf(c_0_76, plain, (X1=X2|X1=X3|~r2_hidden(X1,k1_enumset1(X3,X3,X2))), inference(er,[status(thm)],[c_0_65])).
cnf(c_0_77, plain, (esk5_3(X1,X1,k1_xboole_0)=X1|r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_78, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk11_1(X2)!=k3_mcart_1(X3,X1,X4)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_79, negated_conjecture, (k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=esk4_0|k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=esk4_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_69]), c_0_70])]), c_0_71]), c_0_72]), c_0_73]), c_0_74])).
cnf(c_0_80, negated_conjecture, (k7_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=k2_mcart_1(esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_70]), c_0_74]), c_0_73]), c_0_72])).
cnf(c_0_81, plain, (r2_hidden(esk11_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_82, plain, (esk5_3(X1,X1,k1_xboole_0)=X1|X2=X3|X2=X4), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_83, plain, (X2=k1_xboole_0|esk11_1(X2)!=k4_tarski(k4_tarski(X3,X1),X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_78, c_0_58])).
cnf(c_0_84, negated_conjecture, (k4_tarski(k4_tarski(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),esk4_0),k2_mcart_1(esk4_0))=esk4_0|k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=esk4_0), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_79]), c_0_70])]), c_0_72]), c_0_73]), c_0_74]), c_0_80])).
cnf(c_0_85, plain, (esk11_1(k1_enumset1(X1,X1,X2))=X2|esk11_1(k1_enumset1(X1,X1,X2))=X1|k1_enumset1(X1,X1,X2)=k1_xboole_0), inference(spm,[status(thm)],[c_0_76, c_0_81])).
cnf(c_0_86, plain, (X3=k2_tarski(X1,X2)|esk5_3(X1,X2,X3)!=X2|~r2_hidden(esk5_3(X1,X2,X3),X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_87, plain, (esk5_3(X1,X1,k1_xboole_0)=X1|X2=X3), inference(ef,[status(thm)],[c_0_82])).
cnf(c_0_88, negated_conjecture, (k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=esk4_0|X1=k1_xboole_0|esk11_1(X1)!=esk4_0|~r2_hidden(esk4_0,X1)), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_89, plain, (esk11_1(k1_enumset1(X1,X1,X1))=X1|k1_enumset1(X1,X1,X1)=k1_xboole_0), inference(er,[status(thm)],[inference(ef,[status(thm)],[c_0_85])])).
cnf(c_0_90, plain, (X3=k1_enumset1(X1,X1,X2)|esk5_3(X1,X2,X3)!=X2|~r2_hidden(esk5_3(X1,X2,X3),X3)), inference(rw,[status(thm)],[c_0_86, c_0_28])).
cnf(c_0_91, plain, (esk5_3(X1,X1,k1_xboole_0)=X1), inference(er,[status(thm)],[inference(ef,[status(thm)],[c_0_87])])).
cnf(c_0_92, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk11_1(X2)!=k3_mcart_1(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_93, negated_conjecture, (k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=esk4_0|k1_enumset1(esk4_0,esk4_0,esk4_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89])]), c_0_44])])).
cnf(c_0_94, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|k7_mcart_1(X1,X2,X3,X4)!=X4|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(spm,[status(thm)],[c_0_71, c_0_64])).
cnf(c_0_95, plain, (k1_enumset1(X1,X1,X1)=k1_xboole_0|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_96, plain, (X2=k1_xboole_0|esk11_1(X2)!=k4_tarski(k4_tarski(X1,X3),X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_92, c_0_58])).
cnf(c_0_97, negated_conjecture, (k4_tarski(k4_tarski(esk4_0,k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)),k2_mcart_1(esk4_0))=esk4_0|k1_enumset1(esk4_0,esk4_0,esk4_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_93]), c_0_80]), c_0_70])]), c_0_72]), c_0_73]), c_0_74])).
cnf(c_0_98, negated_conjecture, (k7_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)!=esk4_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_70]), c_0_74]), c_0_73]), c_0_72])).
cnf(c_0_99, plain, (X1=X2|~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_95]), c_0_52])).
cnf(c_0_100, negated_conjecture, (k1_enumset1(esk4_0,esk4_0,esk4_0)=k1_xboole_0|X1=k1_xboole_0|esk11_1(X1)!=esk4_0|~r2_hidden(esk4_0,X1)), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_101, negated_conjecture, (k2_mcart_1(esk4_0)!=esk4_0), inference(rw,[status(thm)],[c_0_98, c_0_80])).
cnf(c_0_102, plain, (k2_mcart_1(X1)=X2|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_99, c_0_60])).
cnf(c_0_103, negated_conjecture, (k1_enumset1(esk4_0,esk4_0,esk4_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_89])]), c_0_44])])).
cnf(c_0_104, negated_conjecture, (~r2_hidden(esk4_0,k1_xboole_0)), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102])])).
cnf(c_0_105, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_103]), c_0_104]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 106
# Proof object clause steps            : 83
# Proof object formula steps           : 23
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 37
# Proof object simplifying inferences  : 62
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 30
# Processed clauses                    : 481
# ...of these trivial                  : 1
# ...subsumed                          : 273
# ...remaining for further processing  : 207
# Other redundant clauses eliminated   : 356
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 34
# Backward-rewritten                   : 7
# Generated clauses                    : 3740
# ...of the previous two non-trivial   : 3271
# Contextual simplify-reflections      : 5
# Paramodulations                      : 3319
# Factorizations                       : 68
# Equation resolutions                 : 356
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 125
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 12
#    Non-unit-clauses                  : 99
# Current number of unprocessed clauses: 2036
# ...number of literals in the above   : 9655
# Current number of archived formulas  : 0
# Current number of archived clauses   : 74
# Clause-clause subsumption calls (NU) : 1578
# Rec. Clause-clause subsumption calls : 863
# Non-unit clause-clause subsumptions  : 94
# Unit Clause-clause subsumption calls : 136
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 11
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 40393

# -------------------------------------------------
# User time                : 0.044 s
# System time              : 0.000 s
# Total time               : 0.044 s
# Maximum resident set size: 28684 pages
