# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t37_mcart_1, conjecture, ![X1, X2, X3, X4, X5, X6]:(k3_zfmisc_1(X1,X2,X3)=k3_zfmisc_1(X4,X5,X6)=>(k3_zfmisc_1(X1,X2,X3)=k1_xboole_0|(X1=X4&(X2=X5&X3=X6)))), file('data/MPTP/proofs/0cde47028424441d89a783120acb6f2b.E_input', t37_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/0cde47028424441d89a783120acb6f2b.E_input', d3_zfmisc_1)).
fof(t138_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=>(k2_zfmisc_1(X1,X2)=k1_xboole_0|(r1_tarski(X1,X3)&r1_tarski(X2,X4)))), file('data/MPTP/proofs/0cde47028424441d89a783120acb6f2b.E_input', t138_zfmisc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/0cde47028424441d89a783120acb6f2b.E_input', d10_xboole_0)).
fof(t117_zfmisc_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&((r1_tarski(k2_zfmisc_1(X2,X1),k2_zfmisc_1(X3,X1))|r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))&~(r1_tarski(X2,X3))))), file('data/MPTP/proofs/0cde47028424441d89a783120acb6f2b.E_input', t117_zfmisc_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/0cde47028424441d89a783120acb6f2b.E_input', t113_zfmisc_1)).
fof(c_0_6, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6]:(k3_zfmisc_1(X1,X2,X3)=k3_zfmisc_1(X4,X5,X6)=>(k3_zfmisc_1(X1,X2,X3)=k1_xboole_0|(X1=X4&(X2=X5&X3=X6))))), inference(assume_negation,[status(cth)],[t37_mcart_1])).
fof(c_0_7, negated_conjecture, (k3_zfmisc_1(esk1_0,esk2_0,esk3_0)=k3_zfmisc_1(esk4_0,esk5_0,esk6_0)&(k3_zfmisc_1(esk1_0,esk2_0,esk3_0)!=k1_xboole_0&(esk1_0!=esk4_0|(esk2_0!=esk5_0|esk3_0!=esk6_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
fof(c_0_8, plain, ![X13, X14, X15]:k3_zfmisc_1(X13,X14,X15)=k2_zfmisc_1(k2_zfmisc_1(X13,X14),X15), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_9, plain, ![X16, X17, X18, X19]:((r1_tarski(X16,X18)|k2_zfmisc_1(X16,X17)=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X16,X17),k2_zfmisc_1(X18,X19)))&(r1_tarski(X17,X19)|k2_zfmisc_1(X16,X17)=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X16,X17),k2_zfmisc_1(X18,X19)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t138_zfmisc_1])])])).
cnf(c_0_10, negated_conjecture, (k3_zfmisc_1(esk1_0,esk2_0,esk3_0)=k3_zfmisc_1(esk4_0,esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, negated_conjecture, (k3_zfmisc_1(esk1_0,esk2_0,esk3_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_13, plain, ![X20, X21]:(((r1_tarski(X20,X21)|X20!=X21)&(r1_tarski(X21,X20)|X20!=X21))&(~r1_tarski(X20,X21)|~r1_tarski(X21,X20)|X20=X21)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_14, plain, (r1_tarski(X1,X2)|k2_zfmisc_1(X3,X1)=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X4,X2))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0),esk6_0)=k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_11]), c_0_11])).
cnf(c_0_16, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_12, c_0_11])).
cnf(c_0_17, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, negated_conjecture, (r1_tarski(esk6_0,X1)|~r1_tarski(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),k2_zfmisc_1(X2,X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])).
cnf(c_0_19, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_17])).
fof(c_0_20, plain, ![X22, X23, X24]:((~r1_tarski(k2_zfmisc_1(X23,X22),k2_zfmisc_1(X24,X22))|r1_tarski(X23,X24)|X22=k1_xboole_0)&(~r1_tarski(k2_zfmisc_1(X22,X23),k2_zfmisc_1(X22,X24))|r1_tarski(X23,X24)|X22=k1_xboole_0)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t117_zfmisc_1])])])])).
cnf(c_0_21, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22, negated_conjecture, (r1_tarski(esk6_0,esk3_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_23, negated_conjecture, (k2_zfmisc_1(X1,X2)=k1_xboole_0|r1_tarski(X2,esk6_0)|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_24, plain, (r1_tarski(X1,X2)|k2_zfmisc_1(X1,X3)=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_25, plain, (r1_tarski(X1,X3)|X2=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, negated_conjecture, (esk6_0=esk3_0|~r1_tarski(esk3_0,esk6_0)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27, negated_conjecture, (r1_tarski(esk3_0,esk6_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_19]), c_0_16])).
cnf(c_0_28, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk4_0,esk5_0),X1)|~r1_tarski(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),k2_zfmisc_1(X1,X2))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_15]), c_0_16])).
cnf(c_0_29, negated_conjecture, (esk6_0=k1_xboole_0|r1_tarski(X1,k2_zfmisc_1(esk4_0,esk5_0))|~r1_tarski(k2_zfmisc_1(X1,esk6_0),k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_25, c_0_15])).
cnf(c_0_30, negated_conjecture, (esk6_0=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_27])])).
cnf(c_0_31, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk4_0,esk5_0),k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_28, c_0_19])).
cnf(c_0_32, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(X1,k2_zfmisc_1(esk4_0,esk5_0))|~r1_tarski(k2_zfmisc_1(X1,esk3_0),k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_30]), c_0_30])).
cnf(c_0_33, negated_conjecture, (k2_zfmisc_1(esk4_0,esk5_0)=k1_xboole_0|r1_tarski(esk4_0,esk1_0)), inference(spm,[status(thm)],[c_0_24, c_0_31])).
cnf(c_0_34, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk4_0,esk5_0))), inference(spm,[status(thm)],[c_0_32, c_0_19])).
cnf(c_0_35, negated_conjecture, (k2_zfmisc_1(esk4_0,esk5_0)=k1_xboole_0|esk4_0=esk1_0|~r1_tarski(esk1_0,esk4_0)), inference(spm,[status(thm)],[c_0_21, c_0_33])).
cnf(c_0_36, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k1_xboole_0|esk3_0=k1_xboole_0|r1_tarski(esk1_0,esk4_0)), inference(spm,[status(thm)],[c_0_24, c_0_34])).
fof(c_0_37, plain, ![X25, X26]:((k2_zfmisc_1(X25,X26)!=k1_xboole_0|(X25=k1_xboole_0|X26=k1_xboole_0))&((X25!=k1_xboole_0|k2_zfmisc_1(X25,X26)=k1_xboole_0)&(X26!=k1_xboole_0|k2_zfmisc_1(X25,X26)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_38, negated_conjecture, (k2_zfmisc_1(esk4_0,esk5_0)=k2_zfmisc_1(esk1_0,esk2_0)|esk3_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_34]), c_0_31])])).
cnf(c_0_39, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k1_xboole_0|k2_zfmisc_1(esk4_0,esk5_0)=k1_xboole_0|esk3_0=k1_xboole_0|esk4_0=esk1_0), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_40, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_41, negated_conjecture, (esk1_0!=esk4_0|esk2_0!=esk5_0|esk3_0!=esk6_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_42, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k1_xboole_0|esk4_0=esk1_0|esk3_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_43, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_44, negated_conjecture, (esk4_0!=esk1_0|esk5_0!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_30])])).
cnf(c_0_45, negated_conjecture, (esk3_0=k1_xboole_0|esk4_0=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_42]), c_0_43])])).
cnf(c_0_46, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k1_xboole_0|esk3_0=k1_xboole_0|r1_tarski(esk2_0,esk5_0)), inference(spm,[status(thm)],[c_0_14, c_0_34])).
cnf(c_0_47, negated_conjecture, (esk3_0=k1_xboole_0|esk5_0!=esk2_0), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_48, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k1_xboole_0|esk3_0=k1_xboole_0|~r1_tarski(esk5_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_46]), c_0_47])).
cnf(c_0_49, negated_conjecture, (k2_zfmisc_1(esk4_0,esk5_0)=k1_xboole_0|r1_tarski(esk5_0,esk2_0)), inference(spm,[status(thm)],[c_0_14, c_0_31])).
cnf(c_0_50, negated_conjecture, (k2_zfmisc_1(esk4_0,esk5_0)=k1_xboole_0|k2_zfmisc_1(esk1_0,esk2_0)=k1_xboole_0|esk3_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_51, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k1_xboole_0|esk3_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_38, c_0_50])).
cnf(c_0_52, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_53, negated_conjecture, (esk3_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_51]), c_0_43])])).
cnf(c_0_54, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_52])).
cnf(c_0_55, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_53]), c_0_54])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 56
# Proof object clause steps            : 43
# Proof object formula steps           : 13
# Proof object conjectures             : 35
# Proof object clause conjectures      : 32
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 23
# Proof object simplifying inferences  : 25
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 14
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 13
# Processed clauses                    : 102
# ...of these trivial                  : 0
# ...subsumed                          : 23
# ...remaining for further processing  : 79
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 9
# Backward-rewritten                   : 32
# Generated clauses                    : 225
# ...of the previous two non-trivial   : 217
# Contextual simplify-reflections      : 1
# Paramodulations                      : 221
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 22
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 17
# Current number of unprocessed clauses: 124
# ...number of literals in the above   : 472
# Current number of archived formulas  : 0
# Current number of archived clauses   : 54
# Clause-clause subsumption calls (NU) : 417
# Rec. Clause-clause subsumption calls : 397
# Non-unit clause-clause subsumptions  : 31
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3960

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.003 s
# Total time               : 0.011 s
# Maximum resident set size: 28740 pages
