# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.006 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t69_xboole_1, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>~((r1_tarski(X2,X1)&r1_xboole_0(X2,X1)))), file('data/MPTP/proofs/0d1ce84377854bc2a4b1ed21fa205d7c.E_input', t69_xboole_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/0d1ce84377854bc2a4b1ed21fa205d7c.E_input', t3_subset)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/0d1ce84377854bc2a4b1ed21fa205d7c.E_input', d1_xboole_0)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/0d1ce84377854bc2a4b1ed21fa205d7c.E_input', t3_xboole_0)).
fof(t17_yellow19, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_waybel_9(X1,k3_yellow19(X1,k2_struct_0(X1),X2),X3)<=>r1_waybel_7(X1,X2,X3))))), file('data/MPTP/proofs/0d1ce84377854bc2a4b1ed21fa205d7c.E_input', t17_yellow19)).
fof(dt_k3_yellow19, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(~(v1_xboole_0(X3))&(v2_waybel_0(X3,k3_yellow_1(X2))&(v13_waybel_0(X3,k3_yellow_1(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))))))=>(~(v2_struct_0(k3_yellow19(X1,X2,X3)))&(v6_waybel_0(k3_yellow19(X1,X2,X3),X1)&l1_waybel_0(k3_yellow19(X1,X2,X3),X1)))), file('data/MPTP/proofs/0d1ce84377854bc2a4b1ed21fa205d7c.E_input', dt_k3_yellow19)).
fof(d2_yellow_1, axiom, ![X1]:k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1)), file('data/MPTP/proofs/0d1ce84377854bc2a4b1ed21fa205d7c.E_input', d2_yellow_1)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/0d1ce84377854bc2a4b1ed21fa205d7c.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/0d1ce84377854bc2a4b1ed21fa205d7c.E_input', d4_subset_1)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/0d1ce84377854bc2a4b1ed21fa205d7c.E_input', dt_l1_pre_topc)).
fof(rc7_pre_topc, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>?[X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))&(~(v1_xboole_0(X2))&v4_pre_topc(X2,X1)))), file('data/MPTP/proofs/0d1ce84377854bc2a4b1ed21fa205d7c.E_input', rc7_pre_topc)).
fof(fc4_yellow19, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(~(v1_xboole_0(X3))&(v2_waybel_0(X3,k3_yellow_1(X2))&(v13_waybel_0(X3,k3_yellow_1(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))))))=>(~(v2_struct_0(k3_yellow19(X1,X2,X3)))&(v3_orders_2(k3_yellow19(X1,X2,X3))&(v4_orders_2(k3_yellow19(X1,X2,X3))&v6_waybel_0(k3_yellow19(X1,X2,X3),X1))))), file('data/MPTP/proofs/0d1ce84377854bc2a4b1ed21fa205d7c.E_input', fc4_yellow19)).
fof(t15_yellow19, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2)))), file('data/MPTP/proofs/0d1ce84377854bc2a4b1ed21fa205d7c.E_input', t15_yellow19)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/0d1ce84377854bc2a4b1ed21fa205d7c.E_input', d3_struct_0)).
fof(t12_yellow19, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_waybel_9(X1,X2,X3)<=>r1_waybel_7(X1,k2_yellow19(X1,X2),X3))))), file('data/MPTP/proofs/0d1ce84377854bc2a4b1ed21fa205d7c.E_input', t12_yellow19)).
fof(fc5_yellow19, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(~(v1_xboole_0(X3))&(v1_subset_1(X3,u1_struct_0(k3_yellow_1(X2)))&(v2_waybel_0(X3,k3_yellow_1(X2))&(v13_waybel_0(X3,k3_yellow_1(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))))))))=>(~(v2_struct_0(k3_yellow19(X1,X2,X3)))&(v6_waybel_0(k3_yellow19(X1,X2,X3),X1)&v7_waybel_0(k3_yellow19(X1,X2,X3))))), file('data/MPTP/proofs/0d1ce84377854bc2a4b1ed21fa205d7c.E_input', fc5_yellow19)).
fof(c_0_16, plain, ![X33, X34]:(v1_xboole_0(X34)|(~r1_tarski(X34,X33)|~r1_xboole_0(X34,X33))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t69_xboole_1])])])).
fof(c_0_17, plain, ![X14, X15]:((~m1_subset_1(X14,k1_zfmisc_1(X15))|r1_tarski(X14,X15))&(~r1_tarski(X14,X15)|m1_subset_1(X14,k1_zfmisc_1(X15)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_18, plain, ![X19, X20, X21]:((~v1_xboole_0(X19)|~r2_hidden(X20,X19))&(r2_hidden(esk4_1(X21),X21)|v1_xboole_0(X21))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_19, plain, ![X36, X37, X39, X40, X41]:(((r2_hidden(esk6_2(X36,X37),X36)|r1_xboole_0(X36,X37))&(r2_hidden(esk6_2(X36,X37),X37)|r1_xboole_0(X36,X37)))&(~r2_hidden(X41,X39)|~r2_hidden(X41,X40)|~r1_xboole_0(X39,X40))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_20, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_waybel_9(X1,k3_yellow19(X1,k2_struct_0(X1),X2),X3)<=>r1_waybel_7(X1,X2,X3)))))), inference(assume_negation,[status(cth)],[t17_yellow19])).
cnf(c_0_21, plain, (v1_xboole_0(X1)|~r1_tarski(X1,X2)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (r2_hidden(esk6_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_25, plain, ![X8, X9, X10]:((~v2_struct_0(k3_yellow19(X8,X9,X10))|(v2_struct_0(X8)|~l1_struct_0(X8)|(v1_xboole_0(X9)|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X8)))|(v1_xboole_0(X10)|(~v2_waybel_0(X10,k3_yellow_1(X9))|(~v13_waybel_0(X10,k3_yellow_1(X9))|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X9))))))))))&((v6_waybel_0(k3_yellow19(X8,X9,X10),X8)|(v2_struct_0(X8)|~l1_struct_0(X8)|(v1_xboole_0(X9)|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X8)))|(v1_xboole_0(X10)|(~v2_waybel_0(X10,k3_yellow_1(X9))|(~v13_waybel_0(X10,k3_yellow_1(X9))|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X9))))))))))&(l1_waybel_0(k3_yellow19(X8,X9,X10),X8)|(v2_struct_0(X8)|~l1_struct_0(X8)|(v1_xboole_0(X9)|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X8)))|(v1_xboole_0(X10)|(~v2_waybel_0(X10,k3_yellow_1(X9))|(~v13_waybel_0(X10,k3_yellow_1(X9))|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X9)))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_yellow19])])])])).
fof(c_0_26, plain, ![X7]:k3_yellow_1(X7)=k3_lattice3(k1_lattice3(X7)), inference(variable_rename,[status(thm)],[d2_yellow_1])).
fof(c_0_27, plain, ![X32]:m1_subset_1(k2_subset_1(X32),k1_zfmisc_1(X32)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_28, plain, ![X11]:k2_subset_1(X11)=X11, inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_29, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v1_xboole_0(esk2_0)&(v1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))&(v2_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))&(v13_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))))))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&((~r3_waybel_9(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk3_0)|~r1_waybel_7(esk1_0,esk2_0,esk3_0))&(r3_waybel_9(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk3_0)|r1_waybel_7(esk1_0,esk2_0,esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_20])])])])).
fof(c_0_30, plain, ![X31]:(~l1_pre_topc(X31)|l1_struct_0(X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_31, plain, (v1_xboole_0(X1)|~r1_xboole_0(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_32, plain, (r1_xboole_0(X1,X2)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_33, plain, ![X29]:((m1_subset_1(esk5_1(X29),k1_zfmisc_1(u1_struct_0(X29)))|(v2_struct_0(X29)|(~v2_pre_topc(X29)|~l1_pre_topc(X29))))&((~v1_xboole_0(esk5_1(X29))|(v2_struct_0(X29)|(~v2_pre_topc(X29)|~l1_pre_topc(X29))))&(v4_pre_topc(esk5_1(X29),X29)|(v2_struct_0(X29)|(~v2_pre_topc(X29)|~l1_pre_topc(X29)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[rc7_pre_topc])])])])])).
fof(c_0_34, plain, ![X23, X24, X25]:((~v2_struct_0(k3_yellow19(X23,X24,X25))|(v2_struct_0(X23)|~l1_struct_0(X23)|(v1_xboole_0(X24)|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X23)))|(v1_xboole_0(X25)|(~v2_waybel_0(X25,k3_yellow_1(X24))|(~v13_waybel_0(X25,k3_yellow_1(X24))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X24))))))))))&((v3_orders_2(k3_yellow19(X23,X24,X25))|(v2_struct_0(X23)|~l1_struct_0(X23)|(v1_xboole_0(X24)|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X23)))|(v1_xboole_0(X25)|(~v2_waybel_0(X25,k3_yellow_1(X24))|(~v13_waybel_0(X25,k3_yellow_1(X24))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X24))))))))))&((v4_orders_2(k3_yellow19(X23,X24,X25))|(v2_struct_0(X23)|~l1_struct_0(X23)|(v1_xboole_0(X24)|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X23)))|(v1_xboole_0(X25)|(~v2_waybel_0(X25,k3_yellow_1(X24))|(~v13_waybel_0(X25,k3_yellow_1(X24))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X24))))))))))&(v6_waybel_0(k3_yellow19(X23,X24,X25),X23)|(v2_struct_0(X23)|~l1_struct_0(X23)|(v1_xboole_0(X24)|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X23)))|(v1_xboole_0(X25)|(~v2_waybel_0(X25,k3_yellow_1(X24))|(~v13_waybel_0(X25,k3_yellow_1(X24))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X24))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc4_yellow19])])])])).
fof(c_0_35, plain, ![X12, X13]:(v2_struct_0(X12)|~l1_struct_0(X12)|(v1_xboole_0(X13)|(~v1_subset_1(X13,u1_struct_0(k3_yellow_1(k2_struct_0(X12))))|(~v2_waybel_0(X13,k3_yellow_1(k2_struct_0(X12)))|(~v13_waybel_0(X13,k3_yellow_1(k2_struct_0(X12)))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X12))))))))|X13=k2_yellow19(X12,k3_yellow19(X12,k2_struct_0(X12),X13)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t15_yellow19])])])])).
cnf(c_0_36, plain, (l1_waybel_0(k3_yellow19(X1,X2,X3),X1)|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_37, plain, (k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_38, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_39, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_40, negated_conjecture, (v13_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_41, plain, ![X35]:(~l1_struct_0(X35)|k2_struct_0(X35)=u1_struct_0(X35)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_42, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_43, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_44, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_45, negated_conjecture, (v2_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_46, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_47, plain, (m1_subset_1(esk5_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_48, plain, (v2_struct_0(X1)|~v1_xboole_0(esk5_1(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_49, plain, (v4_orders_2(k3_yellow19(X1,X2,X3))|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_50, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2))|~l1_struct_0(X1)|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_51, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_52, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~v2_struct_0(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_53, plain, ![X26, X27, X28]:((~r3_waybel_9(X26,X27,X28)|r1_waybel_7(X26,k2_yellow19(X26,X27),X28)|~m1_subset_1(X28,u1_struct_0(X26))|(v2_struct_0(X27)|(~v4_orders_2(X27)|(~v7_waybel_0(X27)|~l1_waybel_0(X27,X26))))|(v2_struct_0(X26)|(~v2_pre_topc(X26)|~l1_pre_topc(X26))))&(~r1_waybel_7(X26,k2_yellow19(X26,X27),X28)|r3_waybel_9(X26,X27,X28)|~m1_subset_1(X28,u1_struct_0(X26))|(v2_struct_0(X27)|(~v4_orders_2(X27)|(~v7_waybel_0(X27)|~l1_waybel_0(X27,X26))))|(v2_struct_0(X26)|(~v2_pre_topc(X26)|~l1_pre_topc(X26))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t12_yellow19])])])])])).
cnf(c_0_54, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|l1_waybel_0(k3_yellow19(X1,X2,X3),X1)|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_37]), c_0_37]), c_0_37])).
cnf(c_0_55, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_56, negated_conjecture, (v13_waybel_0(esk2_0,k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))), inference(rw,[status(thm)],[c_0_40, c_0_37])).
cnf(c_0_57, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_58, negated_conjecture, (l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_59, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))))), inference(rw,[status(thm)],[c_0_44, c_0_37])).
cnf(c_0_60, negated_conjecture, (v2_waybel_0(esk2_0,k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))), inference(rw,[status(thm)],[c_0_45, c_0_37])).
cnf(c_0_61, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])).
cnf(c_0_62, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_63, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_64, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|v4_orders_2(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_37]), c_0_37]), c_0_37])).
cnf(c_0_65, plain, (X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2))|v2_struct_0(X1)|v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~v1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(X1)))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(X1))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_37]), c_0_37]), c_0_37]), c_0_37])).
cnf(c_0_66, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(esk1_0)))))), inference(rw,[status(thm)],[c_0_51, c_0_37])).
cnf(c_0_67, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_68, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_struct_0(k3_yellow19(X1,X2,X3))|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_37]), c_0_37]), c_0_37])).
fof(c_0_69, plain, ![X16, X17, X18]:((~v2_struct_0(k3_yellow19(X16,X17,X18))|(v2_struct_0(X16)|~l1_struct_0(X16)|(v1_xboole_0(X17)|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|(v1_xboole_0(X18)|(~v1_subset_1(X18,u1_struct_0(k3_yellow_1(X17)))|(~v2_waybel_0(X18,k3_yellow_1(X17))|(~v13_waybel_0(X18,k3_yellow_1(X17))|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X17)))))))))))&((v6_waybel_0(k3_yellow19(X16,X17,X18),X16)|(v2_struct_0(X16)|~l1_struct_0(X16)|(v1_xboole_0(X17)|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|(v1_xboole_0(X18)|(~v1_subset_1(X18,u1_struct_0(k3_yellow_1(X17)))|(~v2_waybel_0(X18,k3_yellow_1(X17))|(~v13_waybel_0(X18,k3_yellow_1(X17))|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X17)))))))))))&(v7_waybel_0(k3_yellow19(X16,X17,X18))|(v2_struct_0(X16)|~l1_struct_0(X16)|(v1_xboole_0(X17)|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|(v1_xboole_0(X18)|(~v1_subset_1(X18,u1_struct_0(k3_yellow_1(X17)))|(~v2_waybel_0(X18,k3_yellow_1(X17))|(~v13_waybel_0(X18,k3_yellow_1(X17))|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X17))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc5_yellow19])])])])).
cnf(c_0_70, plain, (r3_waybel_9(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_waybel_7(X1,k2_yellow19(X1,X2),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_71, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_72, plain, (l1_waybel_0(k3_yellow19(X1,u1_struct_0(X1),X2),X1)|v1_xboole_0(u1_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(X1))))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_73, negated_conjecture, (v13_waybel_0(esk2_0,k3_lattice3(k1_lattice3(u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58])])).
cnf(c_0_74, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(esk1_0))))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_57]), c_0_58])])).
cnf(c_0_75, negated_conjecture, (v2_waybel_0(esk2_0,k3_lattice3(k1_lattice3(u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_57]), c_0_58])])).
cnf(c_0_76, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_43])]), c_0_63])).
cnf(c_0_77, plain, (v4_orders_2(k3_yellow19(X1,u1_struct_0(X1),X2))|v1_xboole_0(u1_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(X1))))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))), inference(spm,[status(thm)],[c_0_64, c_0_55])).
cnf(c_0_78, negated_conjecture, (k2_yellow19(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0))=esk2_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_56]), c_0_58]), c_0_59]), c_0_60]), c_0_66])]), c_0_67]), c_0_63])).
cnf(c_0_79, plain, (v1_xboole_0(u1_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(X1))))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))|~v2_struct_0(k3_yellow19(X1,u1_struct_0(X1),X2))), inference(spm,[status(thm)],[c_0_68, c_0_55])).
cnf(c_0_80, negated_conjecture, (r3_waybel_9(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk3_0)|r1_waybel_7(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_81, plain, (r1_waybel_7(X1,k2_yellow19(X1,X2),X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_82, plain, (v7_waybel_0(k3_yellow19(X1,X2,X3))|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_subset_1(X3,u1_struct_0(k3_yellow_1(X2)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_83, negated_conjecture, (~r3_waybel_9(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk3_0)|~r1_waybel_7(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_84, negated_conjecture, (r3_waybel_9(esk1_0,X1,esk3_0)|v2_struct_0(X1)|~v4_orders_2(X1)|~v7_waybel_0(X1)|~l1_waybel_0(X1,esk1_0)|~r1_waybel_7(esk1_0,k2_yellow19(esk1_0,X1),esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_43]), c_0_62])]), c_0_63])).
cnf(c_0_85, negated_conjecture, (l1_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_58]), c_0_74]), c_0_75])]), c_0_76]), c_0_67]), c_0_63])).
cnf(c_0_86, negated_conjecture, (v4_orders_2(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_73]), c_0_58]), c_0_74]), c_0_75])]), c_0_76]), c_0_67]), c_0_63])).
cnf(c_0_87, negated_conjecture, (k2_yellow19(esk1_0,k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_57]), c_0_58])])).
cnf(c_0_88, negated_conjecture, (~v2_struct_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_73]), c_0_58]), c_0_74]), c_0_75])]), c_0_76]), c_0_67]), c_0_63])).
cnf(c_0_89, negated_conjecture, (r1_waybel_7(esk1_0,esk2_0,esk3_0)|r3_waybel_9(esk1_0,k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_57]), c_0_58])])).
cnf(c_0_90, negated_conjecture, (r1_waybel_7(esk1_0,k2_yellow19(esk1_0,X1),esk3_0)|v2_struct_0(X1)|~v4_orders_2(X1)|~v7_waybel_0(X1)|~l1_waybel_0(X1,esk1_0)|~r3_waybel_9(esk1_0,X1,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_71]), c_0_43]), c_0_62])]), c_0_63])).
cnf(c_0_91, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|v7_waybel_0(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_subset_1(X3,u1_struct_0(k3_lattice3(k1_lattice3(X2))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_37]), c_0_37]), c_0_37]), c_0_37])).
cnf(c_0_92, negated_conjecture, (~r1_waybel_7(esk1_0,esk2_0,esk3_0)|~r3_waybel_9(esk1_0,k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_57]), c_0_58])])).
cnf(c_0_93, negated_conjecture, (r3_waybel_9(esk1_0,k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0),esk3_0)|~v7_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86]), c_0_87])]), c_0_88]), c_0_89])).
cnf(c_0_94, negated_conjecture, (r1_waybel_7(esk1_0,esk2_0,esk3_0)|~v7_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_85]), c_0_87]), c_0_86])]), c_0_88]), c_0_89])).
cnf(c_0_95, plain, (v7_waybel_0(k3_yellow19(X1,u1_struct_0(X1),X2))|v1_xboole_0(u1_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(X1))))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))|~v1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(X1)))))), inference(spm,[status(thm)],[c_0_91, c_0_55])).
cnf(c_0_96, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(esk1_0)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_57]), c_0_58])])).
cnf(c_0_97, negated_conjecture, (~v7_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_94])).
cnf(c_0_98, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_73]), c_0_58]), c_0_74]), c_0_75]), c_0_96])]), c_0_97]), c_0_76]), c_0_67]), c_0_63]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 99
# Proof object clause steps            : 66
# Proof object formula steps           : 33
# Proof object conjectures             : 37
# Proof object clause conjectures      : 34
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 29
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 96
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 38
# Processed clauses                    : 132
# ...of these trivial                  : 0
# ...subsumed                          : 13
# ...remaining for further processing  : 119
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 69
# ...of the previous two non-trivial   : 61
# Contextual simplify-reflections      : 14
# Paramodulations                      : 69
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 85
#    Positive orientable unit clauses  : 19
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 59
# Current number of unprocessed clauses: 1
# ...number of literals in the above   : 11
# Current number of archived formulas  : 0
# Current number of archived clauses   : 36
# Clause-clause subsumption calls (NU) : 6862
# Rec. Clause-clause subsumption calls : 337
# Non-unit clause-clause subsumptions  : 26
# Unit Clause-clause subsumption calls : 51
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7634

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28728 pages
