# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t17_lattice3, conjecture, ![X1]:((v3_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k13_lattice3(X1,k12_lattice3(X1,X2,X3),X3)=X3))), file('data/MPTP/proofs/0d38bca89c7743e59f0d49741ba49eb2.E_input', t17_lattice3)).
fof(cc1_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v1_lattice3(X1)=>~(v2_struct_0(X1)))), file('data/MPTP/proofs/0d38bca89c7743e59f0d49741ba49eb2.E_input', cc1_lattice3)).
fof(redefinition_k12_lattice3, axiom, ![X1, X2, X3]:(((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)), file('data/MPTP/proofs/0d38bca89c7743e59f0d49741ba49eb2.E_input', redefinition_k12_lattice3)).
fof(commutativity_k12_lattice3, axiom, ![X1, X2, X3]:(((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k12_lattice3(X1,X2,X3)=k12_lattice3(X1,X3,X2)), file('data/MPTP/proofs/0d38bca89c7743e59f0d49741ba49eb2.E_input', commutativity_k12_lattice3)).
fof(l26_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k10_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))))))), file('data/MPTP/proofs/0d38bca89c7743e59f0d49741ba49eb2.E_input', l26_lattice3)).
fof(reflexivity_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>r3_orders_2(X1,X2,X2)), file('data/MPTP/proofs/0d38bca89c7743e59f0d49741ba49eb2.E_input', reflexivity_r3_orders_2)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/0d38bca89c7743e59f0d49741ba49eb2.E_input', redefinition_r3_orders_2)).
fof(redefinition_k13_lattice3, axiom, ![X1, X2, X3]:(((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k13_lattice3(X1,X2,X3)=k10_lattice3(X1,X2,X3)), file('data/MPTP/proofs/0d38bca89c7743e59f0d49741ba49eb2.E_input', redefinition_k13_lattice3)).
fof(l28_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k11_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))))))), file('data/MPTP/proofs/0d38bca89c7743e59f0d49741ba49eb2.E_input', l28_lattice3)).
fof(cc2_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v2_lattice3(X1)=>~(v2_struct_0(X1)))), file('data/MPTP/proofs/0d38bca89c7743e59f0d49741ba49eb2.E_input', cc2_lattice3)).
fof(dt_k11_lattice3, axiom, ![X1, X2, X3]:((l1_orders_2(X1)&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k11_lattice3(X1,X2,X3),u1_struct_0(X1))), file('data/MPTP/proofs/0d38bca89c7743e59f0d49741ba49eb2.E_input', dt_k11_lattice3)).
fof(c_0_11, negated_conjecture, ~(![X1]:((v3_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k13_lattice3(X1,k12_lattice3(X1,X2,X3),X3)=X3)))), inference(assume_negation,[status(cth)],[t17_lattice3])).
fof(c_0_12, plain, ![X9]:(~l1_orders_2(X9)|(~v1_lattice3(X9)|~v2_struct_0(X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattice3])])])).
fof(c_0_13, negated_conjecture, ((v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_lattice3(esk1_0)&(v2_lattice3(esk1_0)&l1_orders_2(esk1_0)))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&k13_lattice3(esk1_0,k12_lattice3(esk1_0,esk2_0,esk3_0),esk3_0)!=esk3_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
cnf(c_0_14, plain, (~l1_orders_2(X1)|~v1_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15, negated_conjecture, (v1_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_16, plain, ![X10, X11, X12]:(~v5_orders_2(X10)|(~v2_lattice3(X10)|~l1_orders_2(X10))|(~m1_subset_1(X11,u1_struct_0(X10))|~m1_subset_1(X12,u1_struct_0(X10)))|k12_lattice3(X10,X11,X12)=k11_lattice3(X10,X11,X12)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k12_lattice3])])).
fof(c_0_17, plain, ![X13, X14, X15]:(~v5_orders_2(X13)|(~v2_lattice3(X13)|~l1_orders_2(X13))|(~m1_subset_1(X14,u1_struct_0(X13))|~m1_subset_1(X15,u1_struct_0(X13)))|k12_lattice3(X13,X14,X15)=k12_lattice3(X13,X15,X14)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k12_lattice3])])).
fof(c_0_18, plain, ![X16, X17, X18, X19, X20]:(((r1_orders_2(X16,X17,X19)|X19!=k10_lattice3(X16,X17,X18)|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v5_orders_2(X16)|(~v1_lattice3(X16)|~l1_orders_2(X16)))))&((r1_orders_2(X16,X18,X19)|X19!=k10_lattice3(X16,X17,X18)|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v5_orders_2(X16)|(~v1_lattice3(X16)|~l1_orders_2(X16)))))&(~m1_subset_1(X20,u1_struct_0(X16))|(~r1_orders_2(X16,X17,X20)|~r1_orders_2(X16,X18,X20)|r1_orders_2(X16,X19,X20))|X19!=k10_lattice3(X16,X17,X18)|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v5_orders_2(X16)|(~v1_lattice3(X16)|~l1_orders_2(X16)))))))&((m1_subset_1(esk4_4(X16,X17,X18,X19),u1_struct_0(X16))|~r1_orders_2(X16,X18,X19)|~r1_orders_2(X16,X17,X19)|X19=k10_lattice3(X16,X17,X18)|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v5_orders_2(X16)|(~v1_lattice3(X16)|~l1_orders_2(X16)))))&(((r1_orders_2(X16,X17,esk4_4(X16,X17,X18,X19))|~r1_orders_2(X16,X18,X19)|~r1_orders_2(X16,X17,X19)|X19=k10_lattice3(X16,X17,X18)|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v5_orders_2(X16)|(~v1_lattice3(X16)|~l1_orders_2(X16)))))&(r1_orders_2(X16,X18,esk4_4(X16,X17,X18,X19))|~r1_orders_2(X16,X18,X19)|~r1_orders_2(X16,X17,X19)|X19=k10_lattice3(X16,X17,X18)|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v5_orders_2(X16)|(~v1_lattice3(X16)|~l1_orders_2(X16))))))&(~r1_orders_2(X16,X19,esk4_4(X16,X17,X18,X19))|~r1_orders_2(X16,X18,X19)|~r1_orders_2(X16,X17,X19)|X19=k10_lattice3(X16,X17,X18)|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v5_orders_2(X16)|(~v1_lattice3(X16)|~l1_orders_2(X16)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l26_lattice3])])])])])])).
fof(c_0_19, plain, ![X22, X23, X24]:(v2_struct_0(X22)|(~v3_orders_2(X22)|~l1_orders_2(X22))|(~m1_subset_1(X23,u1_struct_0(X22))|~m1_subset_1(X24,u1_struct_0(X22)))|r3_orders_2(X22,X23,X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r3_orders_2])])])).
cnf(c_0_20, negated_conjecture, (~v2_struct_0(esk1_0)|~l1_orders_2(esk1_0)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_21, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22, plain, (k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (k12_lattice3(X1,X2,X3)=k12_lattice3(X1,X3,X2)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (X2=k10_lattice3(X1,X3,X4)|v2_struct_0(X1)|~r1_orders_2(X1,X2,esk4_4(X1,X3,X4,X2))|~r1_orders_2(X1,X4,X2)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (r1_orders_2(X1,X2,esk4_4(X1,X3,X2,X4))|X4=k10_lattice3(X1,X3,X2)|v2_struct_0(X1)|~r1_orders_2(X1,X2,X4)|~r1_orders_2(X1,X3,X4)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_26, plain, ![X25, X26, X27]:((~r3_orders_2(X25,X26,X27)|r1_orders_2(X25,X26,X27)|(v2_struct_0(X25)|(~v3_orders_2(X25)|~l1_orders_2(X25))|(~m1_subset_1(X26,u1_struct_0(X25))|~m1_subset_1(X27,u1_struct_0(X25)))))&(~r1_orders_2(X25,X26,X27)|r3_orders_2(X25,X26,X27)|(v2_struct_0(X25)|(~v3_orders_2(X25)|~l1_orders_2(X25))|(~m1_subset_1(X26,u1_struct_0(X25))|~m1_subset_1(X27,u1_struct_0(X25)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_27, plain, (v2_struct_0(X1)|r3_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_29, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_30, negated_conjecture, (~v2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_21])])).
cnf(c_0_31, negated_conjecture, (k13_lattice3(esk1_0,k12_lattice3(esk1_0,esk2_0,esk3_0),esk3_0)!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_32, plain, (k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_33, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_34, negated_conjecture, (v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_35, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_36, plain, ![X28, X29, X30]:(~v5_orders_2(X28)|(~v1_lattice3(X28)|~l1_orders_2(X28))|(~m1_subset_1(X29,u1_struct_0(X28))|~m1_subset_1(X30,u1_struct_0(X28)))|k13_lattice3(X28,X29,X30)=k10_lattice3(X28,X29,X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k13_lattice3])])).
cnf(c_0_37, plain, (X1=k10_lattice3(X2,X3,X4)|~r1_orders_2(X2,X1,esk4_4(X2,X3,X4,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_24, c_0_14])).
cnf(c_0_38, plain, (X1=k10_lattice3(X2,X3,X4)|r1_orders_2(X2,X4,esk4_4(X2,X3,X4,X1))|~r1_orders_2(X2,X3,X1)|~r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_25, c_0_14])).
cnf(c_0_39, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_40, negated_conjecture, (r3_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_21]), c_0_29])]), c_0_30])).
cnf(c_0_41, negated_conjecture, (k13_lattice3(esk1_0,k11_lattice3(esk1_0,esk3_0,esk2_0),esk3_0)!=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33]), c_0_28]), c_0_21]), c_0_34]), c_0_35])])).
cnf(c_0_42, plain, (k13_lattice3(X1,X2,X3)=k10_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_43, plain, (k10_lattice3(X1,X2,X3)=X3|~r1_orders_2(X1,X3,X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_44, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_21]), c_0_29])]), c_0_30])).
fof(c_0_45, plain, ![X31, X32, X33, X34, X35]:(((r1_orders_2(X31,X34,X32)|X34!=k11_lattice3(X31,X32,X33)|~m1_subset_1(X34,u1_struct_0(X31))|~m1_subset_1(X33,u1_struct_0(X31))|~m1_subset_1(X32,u1_struct_0(X31))|(v2_struct_0(X31)|(~v5_orders_2(X31)|(~v2_lattice3(X31)|~l1_orders_2(X31)))))&((r1_orders_2(X31,X34,X33)|X34!=k11_lattice3(X31,X32,X33)|~m1_subset_1(X34,u1_struct_0(X31))|~m1_subset_1(X33,u1_struct_0(X31))|~m1_subset_1(X32,u1_struct_0(X31))|(v2_struct_0(X31)|(~v5_orders_2(X31)|(~v2_lattice3(X31)|~l1_orders_2(X31)))))&(~m1_subset_1(X35,u1_struct_0(X31))|(~r1_orders_2(X31,X35,X32)|~r1_orders_2(X31,X35,X33)|r1_orders_2(X31,X35,X34))|X34!=k11_lattice3(X31,X32,X33)|~m1_subset_1(X34,u1_struct_0(X31))|~m1_subset_1(X33,u1_struct_0(X31))|~m1_subset_1(X32,u1_struct_0(X31))|(v2_struct_0(X31)|(~v5_orders_2(X31)|(~v2_lattice3(X31)|~l1_orders_2(X31)))))))&((m1_subset_1(esk5_4(X31,X32,X33,X34),u1_struct_0(X31))|~r1_orders_2(X31,X34,X33)|~r1_orders_2(X31,X34,X32)|X34=k11_lattice3(X31,X32,X33)|~m1_subset_1(X34,u1_struct_0(X31))|~m1_subset_1(X33,u1_struct_0(X31))|~m1_subset_1(X32,u1_struct_0(X31))|(v2_struct_0(X31)|(~v5_orders_2(X31)|(~v2_lattice3(X31)|~l1_orders_2(X31)))))&(((r1_orders_2(X31,esk5_4(X31,X32,X33,X34),X32)|~r1_orders_2(X31,X34,X33)|~r1_orders_2(X31,X34,X32)|X34=k11_lattice3(X31,X32,X33)|~m1_subset_1(X34,u1_struct_0(X31))|~m1_subset_1(X33,u1_struct_0(X31))|~m1_subset_1(X32,u1_struct_0(X31))|(v2_struct_0(X31)|(~v5_orders_2(X31)|(~v2_lattice3(X31)|~l1_orders_2(X31)))))&(r1_orders_2(X31,esk5_4(X31,X32,X33,X34),X33)|~r1_orders_2(X31,X34,X33)|~r1_orders_2(X31,X34,X32)|X34=k11_lattice3(X31,X32,X33)|~m1_subset_1(X34,u1_struct_0(X31))|~m1_subset_1(X33,u1_struct_0(X31))|~m1_subset_1(X32,u1_struct_0(X31))|(v2_struct_0(X31)|(~v5_orders_2(X31)|(~v2_lattice3(X31)|~l1_orders_2(X31))))))&(~r1_orders_2(X31,esk5_4(X31,X32,X33,X34),X34)|~r1_orders_2(X31,X34,X33)|~r1_orders_2(X31,X34,X32)|X34=k11_lattice3(X31,X32,X33)|~m1_subset_1(X34,u1_struct_0(X31))|~m1_subset_1(X33,u1_struct_0(X31))|~m1_subset_1(X32,u1_struct_0(X31))|(v2_struct_0(X31)|(~v5_orders_2(X31)|(~v2_lattice3(X31)|~l1_orders_2(X31)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l28_lattice3])])])])])])).
fof(c_0_46, plain, ![X37]:(~l1_orders_2(X37)|(~v2_lattice3(X37)|~v2_struct_0(X37))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_lattice3])])])).
fof(c_0_47, plain, ![X38, X39, X40]:(~l1_orders_2(X38)|(~m1_subset_1(X39,u1_struct_0(X38))|~m1_subset_1(X40,u1_struct_0(X38)))|m1_subset_1(k11_lattice3(X38,X39,X40),u1_struct_0(X38))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k11_lattice3])])).
cnf(c_0_48, negated_conjecture, (k10_lattice3(esk1_0,k11_lattice3(esk1_0,esk3_0,esk2_0),esk3_0)!=esk3_0|~m1_subset_1(k11_lattice3(esk1_0,esk3_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_28]), c_0_21]), c_0_15]), c_0_35])])).
cnf(c_0_49, negated_conjecture, (k10_lattice3(esk1_0,X1,X2)=X2|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_21]), c_0_15]), c_0_35])])).
cnf(c_0_50, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|X2!=k11_lattice3(X1,X3,X4)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_51, plain, (~l1_orders_2(X1)|~v2_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_52, plain, (m1_subset_1(k11_lattice3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_53, negated_conjecture, (~r1_orders_2(esk1_0,k11_lattice3(esk1_0,esk3_0,esk2_0),esk3_0)|~m1_subset_1(k11_lattice3(esk1_0,esk3_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_28])])).
cnf(c_0_54, plain, (r1_orders_2(X1,k11_lattice3(X1,X2,X3),X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_50, c_0_51])]), c_0_52])).
cnf(c_0_55, negated_conjecture, (~m1_subset_1(k11_lattice3(esk1_0,esk3_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_33]), c_0_28]), c_0_21]), c_0_34]), c_0_35])])).
cnf(c_0_56, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_52]), c_0_33]), c_0_28]), c_0_21])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 57
# Proof object clause steps            : 34
# Proof object formula steps           : 23
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 42
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 31
# Processed clauses                    : 109
# ...of these trivial                  : 0
# ...subsumed                          : 27
# ...remaining for further processing  : 82
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 63
# ...of the previous two non-trivial   : 60
# Contextual simplify-reflections      : 17
# Paramodulations                      : 57
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 44
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 31
# Current number of unprocessed clauses: 13
# ...number of literals in the above   : 117
# Current number of archived formulas  : 0
# Current number of archived clauses   : 32
# Clause-clause subsumption calls (NU) : 1096
# Rec. Clause-clause subsumption calls : 59
# Non-unit clause-clause subsumptions  : 38
# Unit Clause-clause subsumption calls : 11
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5748

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.003 s
# Total time               : 0.013 s
# Maximum resident set size: 28660 pages
