# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d19_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v12_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r1_orders_2(X1,X4,X3))=>r2_hidden(X4,X2))))))), file('data/MPTP/proofs/0d3c1f185dda4a90a05fb82d5af9c3e4.E_input', d19_waybel_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/0d3c1f185dda4a90a05fb82d5af9c3e4.E_input', t4_subset)).
fof(d15_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k3_waybel_0(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>?[X5]:(m1_subset_1(X5,u1_struct_0(X1))&(r1_orders_2(X1,X4,X5)&r2_hidden(X5,X2))))))))), file('data/MPTP/proofs/0d3c1f185dda4a90a05fb82d5af9c3e4.E_input', d15_waybel_0)).
fof(dt_k3_waybel_0, axiom, ![X1, X2]:((l1_orders_2(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k3_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/0d3c1f185dda4a90a05fb82d5af9c3e4.E_input', dt_k3_waybel_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/0d3c1f185dda4a90a05fb82d5af9c3e4.E_input', d3_tarski)).
fof(t23_waybel_0, conjecture, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v12_waybel_0(X2,X1)<=>r1_tarski(k3_waybel_0(X1,X2),X2)))), file('data/MPTP/proofs/0d3c1f185dda4a90a05fb82d5af9c3e4.E_input', t23_waybel_0)).
fof(c_0_6, plain, ![X19, X20, X21, X22]:((~v12_waybel_0(X20,X19)|(~m1_subset_1(X21,u1_struct_0(X19))|(~m1_subset_1(X22,u1_struct_0(X19))|(~r2_hidden(X21,X20)|~r1_orders_2(X19,X22,X21)|r2_hidden(X22,X20))))|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|~l1_orders_2(X19))&((m1_subset_1(esk4_2(X19,X20),u1_struct_0(X19))|v12_waybel_0(X20,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|~l1_orders_2(X19))&((m1_subset_1(esk5_2(X19,X20),u1_struct_0(X19))|v12_waybel_0(X20,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|~l1_orders_2(X19))&(((r2_hidden(esk4_2(X19,X20),X20)|v12_waybel_0(X20,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|~l1_orders_2(X19))&(r1_orders_2(X19,esk5_2(X19,X20),esk4_2(X19,X20))|v12_waybel_0(X20,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|~l1_orders_2(X19)))&(~r2_hidden(esk5_2(X19,X20),X20)|v12_waybel_0(X20,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|~l1_orders_2(X19)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d19_waybel_0])])])])])).
fof(c_0_7, plain, ![X10, X11, X12]:(~r2_hidden(X10,X11)|~m1_subset_1(X11,k1_zfmisc_1(X12))|m1_subset_1(X10,X12)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_8, plain, ![X25, X26, X27, X28, X30, X32]:((((m1_subset_1(esk6_4(X25,X26,X27,X28),u1_struct_0(X25))|~r2_hidden(X28,X27)|~m1_subset_1(X28,u1_struct_0(X25))|X27!=k3_waybel_0(X25,X26)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X25)))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|~l1_orders_2(X25))&((r1_orders_2(X25,X28,esk6_4(X25,X26,X27,X28))|~r2_hidden(X28,X27)|~m1_subset_1(X28,u1_struct_0(X25))|X27!=k3_waybel_0(X25,X26)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X25)))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|~l1_orders_2(X25))&(r2_hidden(esk6_4(X25,X26,X27,X28),X26)|~r2_hidden(X28,X27)|~m1_subset_1(X28,u1_struct_0(X25))|X27!=k3_waybel_0(X25,X26)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X25)))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|~l1_orders_2(X25))))&(~m1_subset_1(X30,u1_struct_0(X25))|(~r1_orders_2(X25,X28,X30)|~r2_hidden(X30,X26))|r2_hidden(X28,X27)|~m1_subset_1(X28,u1_struct_0(X25))|X27!=k3_waybel_0(X25,X26)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X25)))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|~l1_orders_2(X25)))&((m1_subset_1(esk7_3(X25,X26,X27),u1_struct_0(X25))|X27=k3_waybel_0(X25,X26)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X25)))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|~l1_orders_2(X25))&((~r2_hidden(esk7_3(X25,X26,X27),X27)|(~m1_subset_1(X32,u1_struct_0(X25))|(~r1_orders_2(X25,esk7_3(X25,X26,X27),X32)|~r2_hidden(X32,X26)))|X27=k3_waybel_0(X25,X26)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X25)))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|~l1_orders_2(X25))&((m1_subset_1(esk8_3(X25,X26,X27),u1_struct_0(X25))|r2_hidden(esk7_3(X25,X26,X27),X27)|X27=k3_waybel_0(X25,X26)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X25)))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|~l1_orders_2(X25))&((r1_orders_2(X25,esk7_3(X25,X26,X27),esk8_3(X25,X26,X27))|r2_hidden(esk7_3(X25,X26,X27),X27)|X27=k3_waybel_0(X25,X26)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X25)))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|~l1_orders_2(X25))&(r2_hidden(esk8_3(X25,X26,X27),X26)|r2_hidden(esk7_3(X25,X26,X27),X27)|X27=k3_waybel_0(X25,X26)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X25)))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|~l1_orders_2(X25))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d15_waybel_0])])])])])).
fof(c_0_9, plain, ![X8, X9]:(~l1_orders_2(X8)|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X8)))|m1_subset_1(k3_waybel_0(X8,X9),k1_zfmisc_1(u1_struct_0(X8)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_waybel_0])])).
cnf(c_0_10, plain, (r2_hidden(X4,X1)|~v12_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,X1)|~r1_orders_2(X2,X4,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (r1_orders_2(X1,X2,esk6_4(X1,X3,X4,X2))|~r2_hidden(X2,X4)|~m1_subset_1(X2,u1_struct_0(X1))|X4!=k3_waybel_0(X1,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (m1_subset_1(k3_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (r2_hidden(X1,X2)|~r1_orders_2(X3,X1,X4)|~r2_hidden(X4,X2)|~v12_waybel_0(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)), inference(csr,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15, plain, (r1_orders_2(X1,X2,esk6_4(X1,X3,k3_waybel_0(X1,X3),X2))|~r2_hidden(X2,k3_waybel_0(X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_12, c_0_11])]), c_0_13])).
cnf(c_0_16, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,k3_waybel_0(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_11, c_0_13])).
cnf(c_0_17, plain, (r2_hidden(esk6_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_18, plain, (r2_hidden(X1,X2)|~r2_hidden(esk6_4(X3,X4,k3_waybel_0(X3,X4),X1),X2)|~r2_hidden(X1,k3_waybel_0(X3,X4))|~v12_waybel_0(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~l1_orders_2(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])).
cnf(c_0_19, plain, (r2_hidden(esk6_4(X1,X2,k3_waybel_0(X1,X2),X3),X2)|~r2_hidden(X3,k3_waybel_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_17, c_0_11])]), c_0_13])).
fof(c_0_20, plain, ![X13, X14, X15, X16, X17]:((~r1_tarski(X13,X14)|(~r2_hidden(X15,X13)|r2_hidden(X15,X14)))&((r2_hidden(esk3_2(X16,X17),X16)|r1_tarski(X16,X17))&(~r2_hidden(esk3_2(X16,X17),X17)|r1_tarski(X16,X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_21, negated_conjecture, ~(![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v12_waybel_0(X2,X1)<=>r1_tarski(k3_waybel_0(X1,X2),X2))))), inference(assume_negation,[status(cth)],[t23_waybel_0])).
cnf(c_0_22, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_waybel_0(X3,X2))|~v12_waybel_0(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~l1_orders_2(X3)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_23, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_24, negated_conjecture, (l1_orders_2(esk1_0)&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&((~v12_waybel_0(esk2_0,esk1_0)|~r1_tarski(k3_waybel_0(esk1_0,esk2_0),esk2_0))&(v12_waybel_0(esk2_0,esk1_0)|r1_tarski(k3_waybel_0(esk1_0,esk2_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).
cnf(c_0_25, plain, (r1_tarski(X1,X2)|~r2_hidden(esk3_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (r2_hidden(esk3_2(k3_waybel_0(X1,X2),X3),X2)|r1_tarski(k3_waybel_0(X1,X2),X3)|~v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_27, negated_conjecture, (~v12_waybel_0(esk2_0,esk1_0)|~r1_tarski(k3_waybel_0(esk1_0,esk2_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_28, plain, (r1_tarski(k3_waybel_0(X1,X2),X2)|~v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, negated_conjecture, (v12_waybel_0(esk2_0,esk1_0)|r1_tarski(k3_waybel_0(esk1_0,esk2_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, negated_conjecture, (~v12_waybel_0(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30])])).
cnf(c_0_33, plain, (r2_hidden(X3,X5)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,X3,X1)|~r2_hidden(X1,X4)|~m1_subset_1(X3,u1_struct_0(X2))|X5!=k3_waybel_0(X2,X4)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_34, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_35, negated_conjecture, (r1_tarski(k3_waybel_0(esk1_0,esk2_0),esk2_0)), inference(sr,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_36, plain, (r2_hidden(X1,k3_waybel_0(X2,X3))|~r1_orders_2(X2,X1,X4)|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_33, c_0_11])]), c_0_13])).
cnf(c_0_37, plain, (r1_orders_2(X1,esk5_2(X1,X2),esk4_2(X1,X2))|v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_38, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_39, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,k3_waybel_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40, plain, (r2_hidden(esk5_2(X1,X2),k3_waybel_0(X1,X3))|v12_waybel_0(X2,X1)|~r2_hidden(esk4_2(X1,X2),X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_41, plain, (v12_waybel_0(X2,X1)|~r2_hidden(esk5_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_42, negated_conjecture, (r2_hidden(esk5_2(esk1_0,X1),esk2_0)|v12_waybel_0(X1,esk1_0)|~r2_hidden(esk4_2(esk1_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_29]), c_0_30])])).
cnf(c_0_43, negated_conjecture, (~r2_hidden(esk4_2(esk1_0,esk2_0),esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_29]), c_0_30])]), c_0_32])).
cnf(c_0_44, plain, (r2_hidden(esk4_2(X1,X2),X2)|v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_45, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_29]), c_0_30])]), c_0_32]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 46
# Proof object clause steps            : 33
# Proof object formula steps           : 13
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 27
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 24
# Processed clauses                    : 101
# ...of these trivial                  : 0
# ...subsumed                          : 7
# ...remaining for further processing  : 94
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 1
# Generated clauses                    : 98
# ...of the previous two non-trivial   : 90
# Contextual simplify-reflections      : 14
# Paramodulations                      : 91
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 58
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 52
# Current number of unprocessed clauses: 37
# ...number of literals in the above   : 262
# Current number of archived formulas  : 0
# Current number of archived clauses   : 32
# Clause-clause subsumption calls (NU) : 1688
# Rec. Clause-clause subsumption calls : 263
# Non-unit clause-clause subsumptions  : 24
# Unit Clause-clause subsumption calls : 29
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5746

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.003 s
# Total time               : 0.013 s
# Maximum resident set size: 28656 pages
