# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t45_yellow_6, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v8_pre_topc(X1)&l1_pre_topc(X1))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&(v1_yellow_6(X2,X1)&l1_waybel_0(X2,X1)))))=>k11_yellow_6(X1,X2)=k4_yellow_6(X1,X2))), file('data/MPTP/proofs/0d3ee71f283749258ca38dfe0d907739.E_input', t45_yellow_6)).
fof(d20_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v8_pre_topc(X1)&l1_pre_topc(X1))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&(v3_yellow_6(X2,X1)&l1_waybel_0(X2,X1)))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X3=k11_yellow_6(X1,X2)<=>r2_hidden(X3,k10_yellow_6(X1,X2)))))), file('data/MPTP/proofs/0d3ee71f283749258ca38dfe0d907739.E_input', d20_yellow_6)).
fof(dt_k11_yellow_6, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v8_pre_topc(X1)&l1_pre_topc(X1))))&(~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&(v3_yellow_6(X2,X1)&l1_waybel_0(X2,X1))))))=>m1_subset_1(k11_yellow_6(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/0d3ee71f283749258ca38dfe0d907739.E_input', dt_k11_yellow_6)).
fof(cc4_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(l1_waybel_0(X2,X1)=>((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&v1_yellow_6(X2,X1))))=>(~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&v3_yellow_6(X2,X1))))))), file('data/MPTP/proofs/0d3ee71f283749258ca38dfe0d907739.E_input', cc4_yellow_6)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/0d3ee71f283749258ca38dfe0d907739.E_input', d10_xboole_0)).
fof(t42_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&(v1_yellow_6(X2,X1)&l1_waybel_0(X2,X1)))))=>r2_hidden(k4_yellow_6(X1,X2),k10_yellow_6(X1,X2)))), file('data/MPTP/proofs/0d3ee71f283749258ca38dfe0d907739.E_input', t42_yellow_6)).
fof(l1_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/0d3ee71f283749258ca38dfe0d907739.E_input', l1_zfmisc_1)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/0d3ee71f283749258ca38dfe0d907739.E_input', d1_zfmisc_1)).
fof(fc22_yellow_6, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v8_pre_topc(X1)&l1_pre_topc(X1))))&(~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1)))))=>v1_zfmisc_1(k10_yellow_6(X1,X2))), file('data/MPTP/proofs/0d3ee71f283749258ca38dfe0d907739.E_input', fc22_yellow_6)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/0d3ee71f283749258ca38dfe0d907739.E_input', d1_xboole_0)).
fof(t1_tex_2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((~(v1_xboole_0(X2))&v1_zfmisc_1(X2))=>(r1_tarski(X1,X2)=>X1=X2))), file('data/MPTP/proofs/0d3ee71f283749258ca38dfe0d907739.E_input', t1_tex_2)).
fof(c_0_11, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v8_pre_topc(X1)&l1_pre_topc(X1))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&(v1_yellow_6(X2,X1)&l1_waybel_0(X2,X1)))))=>k11_yellow_6(X1,X2)=k4_yellow_6(X1,X2)))), inference(assume_negation,[status(cth)],[t45_yellow_6])).
fof(c_0_12, plain, ![X14, X15, X16]:((X16!=k11_yellow_6(X14,X15)|r2_hidden(X16,k10_yellow_6(X14,X15))|~m1_subset_1(X16,u1_struct_0(X14))|(v2_struct_0(X15)|(~v4_orders_2(X15)|(~v7_waybel_0(X15)|(~v3_yellow_6(X15,X14)|~l1_waybel_0(X15,X14)))))|(v2_struct_0(X14)|(~v2_pre_topc(X14)|(~v8_pre_topc(X14)|~l1_pre_topc(X14)))))&(~r2_hidden(X16,k10_yellow_6(X14,X15))|X16=k11_yellow_6(X14,X15)|~m1_subset_1(X16,u1_struct_0(X14))|(v2_struct_0(X15)|(~v4_orders_2(X15)|(~v7_waybel_0(X15)|(~v3_yellow_6(X15,X14)|~l1_waybel_0(X15,X14)))))|(v2_struct_0(X14)|(~v2_pre_topc(X14)|(~v8_pre_topc(X14)|~l1_pre_topc(X14)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d20_yellow_6])])])])])).
fof(c_0_13, plain, ![X6, X7]:(v2_struct_0(X6)|(~v2_pre_topc(X6)|(~v8_pre_topc(X6)|~l1_pre_topc(X6)))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|(~v3_yellow_6(X7,X6)|~l1_waybel_0(X7,X6)))))|m1_subset_1(k11_yellow_6(X6,X7),u1_struct_0(X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k11_yellow_6])])])).
fof(c_0_14, plain, ![X10, X11]:((~v2_struct_0(X11)|(v2_struct_0(X11)|(~v4_orders_2(X11)|(~v7_waybel_0(X11)|~v1_yellow_6(X11,X10))))|~l1_waybel_0(X11,X10)|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))&((v4_orders_2(X11)|(v2_struct_0(X11)|(~v4_orders_2(X11)|(~v7_waybel_0(X11)|~v1_yellow_6(X11,X10))))|~l1_waybel_0(X11,X10)|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))&((v7_waybel_0(X11)|(v2_struct_0(X11)|(~v4_orders_2(X11)|(~v7_waybel_0(X11)|~v1_yellow_6(X11,X10))))|~l1_waybel_0(X11,X10)|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))&(v3_yellow_6(X11,X10)|(v2_struct_0(X11)|(~v4_orders_2(X11)|(~v7_waybel_0(X11)|~v1_yellow_6(X11,X10))))|~l1_waybel_0(X11,X10)|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc4_yellow_6])])])])])).
fof(c_0_15, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&(v8_pre_topc(esk1_0)&l1_pre_topc(esk1_0))))&((~v2_struct_0(esk2_0)&(v4_orders_2(esk2_0)&(v7_waybel_0(esk2_0)&(v1_yellow_6(esk2_0,esk1_0)&l1_waybel_0(esk2_0,esk1_0)))))&k11_yellow_6(esk1_0,esk2_0)!=k4_yellow_6(esk1_0,esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])).
fof(c_0_16, plain, ![X17, X18]:(((r1_tarski(X17,X18)|X17!=X18)&(r1_tarski(X18,X17)|X17!=X18))&(~r1_tarski(X17,X18)|~r1_tarski(X18,X17)|X17=X18)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_17, plain, (r2_hidden(X1,k10_yellow_6(X2,X3))|v2_struct_0(X3)|v2_struct_0(X2)|X1!=k11_yellow_6(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~v4_orders_2(X3)|~v7_waybel_0(X3)|~v3_yellow_6(X3,X2)|~l1_waybel_0(X3,X2)|~v2_pre_topc(X2)|~v8_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, plain, (v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(k11_yellow_6(X1,X2),u1_struct_0(X1))|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~v3_yellow_6(X2,X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (v3_yellow_6(X1,X2)|v2_struct_0(X1)|v2_struct_0(X2)|~v4_orders_2(X1)|~v7_waybel_0(X1)|~v1_yellow_6(X1,X2)|~l1_waybel_0(X1,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, negated_conjecture, (v1_yellow_6(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, negated_conjecture, (v7_waybel_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, negated_conjecture, (v4_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_26, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_27, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_28, plain, ![X25, X26]:(v2_struct_0(X25)|(~v2_pre_topc(X25)|~l1_pre_topc(X25))|(v2_struct_0(X26)|(~v4_orders_2(X26)|(~v7_waybel_0(X26)|(~v1_yellow_6(X26,X25)|~l1_waybel_0(X26,X25))))|r2_hidden(k4_yellow_6(X25,X26),k10_yellow_6(X25,X26)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t42_yellow_6])])])])).
fof(c_0_29, plain, ![X8, X9]:((~r1_tarski(k1_tarski(X8),X9)|r2_hidden(X8,X9))&(~r2_hidden(X8,X9)|r1_tarski(k1_tarski(X8),X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])).
cnf(c_0_30, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_31, plain, ![X27, X28, X29, X30, X31, X32]:(((~r2_hidden(X29,X28)|r1_tarski(X29,X27)|X28!=k1_zfmisc_1(X27))&(~r1_tarski(X30,X27)|r2_hidden(X30,X28)|X28!=k1_zfmisc_1(X27)))&((~r2_hidden(esk4_2(X31,X32),X32)|~r1_tarski(esk4_2(X31,X32),X31)|X32=k1_zfmisc_1(X31))&(r2_hidden(esk4_2(X31,X32),X32)|r1_tarski(esk4_2(X31,X32),X31)|X32=k1_zfmisc_1(X31)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
cnf(c_0_32, plain, (r2_hidden(k11_yellow_6(X1,X2),k10_yellow_6(X1,X2))|v2_struct_0(X2)|v2_struct_0(X1)|~v3_yellow_6(X2,X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_17]), c_0_18])).
cnf(c_0_33, negated_conjecture, (v3_yellow_6(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25])]), c_0_26]), c_0_27])).
cnf(c_0_34, negated_conjecture, (v8_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_35, plain, ![X23, X24]:(v2_struct_0(X23)|(~v2_pre_topc(X23)|(~v8_pre_topc(X23)|~l1_pre_topc(X23)))|(v2_struct_0(X24)|(~v4_orders_2(X24)|(~v7_waybel_0(X24)|~l1_waybel_0(X24,X23))))|v1_zfmisc_1(k10_yellow_6(X23,X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc22_yellow_6])])])).
fof(c_0_36, plain, ![X19, X20, X21]:((~v1_xboole_0(X19)|~r2_hidden(X20,X19))&(r2_hidden(esk3_1(X21),X21)|v1_xboole_0(X21))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_37, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(k4_yellow_6(X1,X2),k10_yellow_6(X1,X2))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~v1_yellow_6(X2,X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, plain, (r2_hidden(X1,X2)|~r1_tarski(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_30])).
cnf(c_0_40, plain, (r2_hidden(X1,X3)|~r1_tarski(X1,X2)|X3!=k1_zfmisc_1(X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_41, plain, ![X12, X13]:(v1_xboole_0(X12)|(v1_xboole_0(X13)|~v1_zfmisc_1(X13)|(~r1_tarski(X12,X13)|X12=X13))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t1_tex_2])])])])).
cnf(c_0_42, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_43, negated_conjecture, (r2_hidden(k11_yellow_6(esk1_0,esk2_0),k10_yellow_6(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_20]), c_0_22]), c_0_23]), c_0_24]), c_0_34]), c_0_25])]), c_0_26]), c_0_27])).
cnf(c_0_44, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v1_zfmisc_1(k10_yellow_6(X1,X2))|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_45, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_46, negated_conjecture, (r2_hidden(k4_yellow_6(esk1_0,esk2_0),k10_yellow_6(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25])]), c_0_27]), c_0_26])).
cnf(c_0_47, plain, (r2_hidden(X1,k1_tarski(X1))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_48, plain, (r2_hidden(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_49, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|X1=X2|~v1_zfmisc_1(X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_50, negated_conjecture, (r1_tarski(k1_tarski(k11_yellow_6(esk1_0,esk2_0)),k10_yellow_6(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_51, negated_conjecture, (v1_zfmisc_1(k10_yellow_6(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_20]), c_0_22]), c_0_23]), c_0_24]), c_0_34]), c_0_25])]), c_0_27]), c_0_26])).
cnf(c_0_52, negated_conjecture, (~v1_xboole_0(k10_yellow_6(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_53, plain, (~v1_xboole_0(k1_tarski(X1))), inference(spm,[status(thm)],[c_0_45, c_0_47])).
cnf(c_0_54, plain, (r2_hidden(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_48, c_0_39])).
cnf(c_0_55, negated_conjecture, (r1_tarski(k1_tarski(k4_yellow_6(esk1_0,esk2_0)),k10_yellow_6(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_42, c_0_46])).
cnf(c_0_56, negated_conjecture, (k1_tarski(k11_yellow_6(esk1_0,esk2_0))=k10_yellow_6(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51])]), c_0_52]), c_0_53])).
cnf(c_0_57, plain, (r1_tarski(k1_tarski(X1),k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_42, c_0_54])).
cnf(c_0_58, negated_conjecture, (k1_tarski(k4_yellow_6(esk1_0,esk2_0))=k10_yellow_6(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_55]), c_0_51])]), c_0_52]), c_0_53])).
cnf(c_0_59, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_60, negated_conjecture, (r2_hidden(k11_yellow_6(esk1_0,esk2_0),X1)|~r1_tarski(k10_yellow_6(esk1_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_38, c_0_56])).
cnf(c_0_61, negated_conjecture, (r1_tarski(k10_yellow_6(esk1_0,esk2_0),k1_zfmisc_1(k4_yellow_6(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_62, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(er,[status(thm)],[c_0_59])).
cnf(c_0_63, negated_conjecture, (r2_hidden(k11_yellow_6(esk1_0,esk2_0),k1_zfmisc_1(k4_yellow_6(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_64, negated_conjecture, (r2_hidden(k4_yellow_6(esk1_0,esk2_0),X1)|~r1_tarski(k10_yellow_6(esk1_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_38, c_0_58])).
cnf(c_0_65, negated_conjecture, (r1_tarski(k10_yellow_6(esk1_0,esk2_0),k1_zfmisc_1(k11_yellow_6(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_57, c_0_56])).
cnf(c_0_66, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_67, negated_conjecture, (r1_tarski(k11_yellow_6(esk1_0,esk2_0),k4_yellow_6(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_68, negated_conjecture, (k11_yellow_6(esk1_0,esk2_0)!=k4_yellow_6(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_69, negated_conjecture, (r2_hidden(k4_yellow_6(esk1_0,esk2_0),k1_zfmisc_1(k11_yellow_6(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_70, negated_conjecture, (~r1_tarski(k4_yellow_6(esk1_0,esk2_0),k11_yellow_6(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68])).
cnf(c_0_71, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_69]), c_0_70]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 72
# Proof object clause steps            : 49
# Proof object formula steps           : 23
# Proof object conjectures             : 31
# Proof object clause conjectures      : 28
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 22
# Proof object simplifying inferences  : 48
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 28
# Processed clauses                    : 100
# ...of these trivial                  : 1
# ...subsumed                          : 7
# ...remaining for further processing  : 92
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 107
# ...of the previous two non-trivial   : 86
# Contextual simplify-reflections      : 1
# Paramodulations                      : 101
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 57
#    Positive orientable unit clauses  : 25
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 25
# Current number of unprocessed clauses: 40
# ...number of literals in the above   : 112
# Current number of archived formulas  : 0
# Current number of archived clauses   : 30
# Clause-clause subsumption calls (NU) : 543
# Rec. Clause-clause subsumption calls : 84
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 86
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4263

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28800 pages
