# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t27_lattices, conjecture, ![X1]:((~(v2_struct_0(X1))&(v7_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattices(X1,X2,X3)=>r1_lattices(X1,k2_lattices(X1,X2,X4),k2_lattices(X1,X3,X4))))))), file('data/MPTP/proofs/0d68f3b69a52464f837ec1879c56ced7.E_input', t27_lattices)).
fof(t21_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattices(X1,X2,X3)<=>k2_lattices(X1,X2,X3)=X2)))), file('data/MPTP/proofs/0d68f3b69a52464f837ec1879c56ced7.E_input', t21_lattices)).
fof(dt_k2_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_lattices(X1))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k2_lattices(X1,X2,X3),u1_struct_0(X1))), file('data/MPTP/proofs/0d68f3b69a52464f837ec1879c56ced7.E_input', dt_k2_lattices)).
fof(d7_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v7_lattices(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>k2_lattices(X1,X2,k2_lattices(X1,X3,X4))=k2_lattices(X1,k2_lattices(X1,X2,X3),X4)))))), file('data/MPTP/proofs/0d68f3b69a52464f837ec1879c56ced7.E_input', d7_lattices)).
fof(d8_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>(v8_lattices(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k1_lattices(X1,k2_lattices(X1,X2,X3),X3)=X3)))), file('data/MPTP/proofs/0d68f3b69a52464f837ec1879c56ced7.E_input', d8_lattices)).
fof(d9_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>(v9_lattices(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_lattices(X1,X2,k1_lattices(X1,X2,X3))=X2)))), file('data/MPTP/proofs/0d68f3b69a52464f837ec1879c56ced7.E_input', d9_lattices)).
fof(dt_l3_lattices, axiom, ![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1))), file('data/MPTP/proofs/0d68f3b69a52464f837ec1879c56ced7.E_input', dt_l3_lattices)).
fof(c_0_7, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v7_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattices(X1,X2,X3)=>r1_lattices(X1,k2_lattices(X1,X2,X4),k2_lattices(X1,X3,X4)))))))), inference(assume_negation,[status(cth)],[t27_lattices])).
fof(c_0_8, plain, ![X9, X10, X11]:((~r1_lattices(X9,X10,X11)|k2_lattices(X9,X10,X11)=X10|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(v2_struct_0(X9)|(~v8_lattices(X9)|(~v9_lattices(X9)|~l3_lattices(X9)))))&(k2_lattices(X9,X10,X11)!=X10|r1_lattices(X9,X10,X11)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(v2_struct_0(X9)|(~v8_lattices(X9)|(~v9_lattices(X9)|~l3_lattices(X9)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t21_lattices])])])])])).
fof(c_0_9, negated_conjecture, ((~v2_struct_0(esk1_0)&(v7_lattices(esk1_0)&(v8_lattices(esk1_0)&(v9_lattices(esk1_0)&l3_lattices(esk1_0)))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(m1_subset_1(esk4_0,u1_struct_0(esk1_0))&(r1_lattices(esk1_0,esk2_0,esk3_0)&~r1_lattices(esk1_0,k2_lattices(esk1_0,esk2_0,esk4_0),k2_lattices(esk1_0,esk3_0,esk4_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_7])])])])).
fof(c_0_10, plain, ![X12, X13, X14]:(v2_struct_0(X12)|~l1_lattices(X12)|(~m1_subset_1(X13,u1_struct_0(X12))|~m1_subset_1(X14,u1_struct_0(X12)))|m1_subset_1(k2_lattices(X12,X13,X14),u1_struct_0(X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_lattices])])])).
fof(c_0_11, plain, ![X15, X16, X17, X18]:((~v7_lattices(X15)|(~m1_subset_1(X16,u1_struct_0(X15))|(~m1_subset_1(X17,u1_struct_0(X15))|(~m1_subset_1(X18,u1_struct_0(X15))|k2_lattices(X15,X16,k2_lattices(X15,X17,X18))=k2_lattices(X15,k2_lattices(X15,X16,X17),X18))))|(v2_struct_0(X15)|~l1_lattices(X15)))&((m1_subset_1(esk5_1(X15),u1_struct_0(X15))|v7_lattices(X15)|(v2_struct_0(X15)|~l1_lattices(X15)))&((m1_subset_1(esk6_1(X15),u1_struct_0(X15))|v7_lattices(X15)|(v2_struct_0(X15)|~l1_lattices(X15)))&((m1_subset_1(esk7_1(X15),u1_struct_0(X15))|v7_lattices(X15)|(v2_struct_0(X15)|~l1_lattices(X15)))&(k2_lattices(X15,esk5_1(X15),k2_lattices(X15,esk6_1(X15),esk7_1(X15)))!=k2_lattices(X15,k2_lattices(X15,esk5_1(X15),esk6_1(X15)),esk7_1(X15))|v7_lattices(X15)|(v2_struct_0(X15)|~l1_lattices(X15))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_lattices])])])])])])).
cnf(c_0_12, plain, (k2_lattices(X1,X2,X3)=X2|v2_struct_0(X1)|~r1_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, negated_conjecture, (l3_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, negated_conjecture, (v9_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, negated_conjecture, (v8_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_18, plain, ![X22, X23, X24]:((~v8_lattices(X22)|(~m1_subset_1(X23,u1_struct_0(X22))|(~m1_subset_1(X24,u1_struct_0(X22))|k1_lattices(X22,k2_lattices(X22,X23,X24),X24)=X24))|(v2_struct_0(X22)|~l3_lattices(X22)))&((m1_subset_1(esk8_1(X22),u1_struct_0(X22))|v8_lattices(X22)|(v2_struct_0(X22)|~l3_lattices(X22)))&((m1_subset_1(esk9_1(X22),u1_struct_0(X22))|v8_lattices(X22)|(v2_struct_0(X22)|~l3_lattices(X22)))&(k1_lattices(X22,k2_lattices(X22,esk8_1(X22),esk9_1(X22)),esk9_1(X22))!=esk9_1(X22)|v8_lattices(X22)|(v2_struct_0(X22)|~l3_lattices(X22)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_lattices])])])])])])).
cnf(c_0_19, plain, (v2_struct_0(X1)|m1_subset_1(k2_lattices(X1,X2,X3),u1_struct_0(X1))|~l1_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_20, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_21, plain, (k2_lattices(X1,X2,k2_lattices(X1,X3,X4))=k2_lattices(X1,k2_lattices(X1,X2,X3),X4)|v2_struct_0(X1)|~v7_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_22, negated_conjecture, (v7_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_23, negated_conjecture, (k2_lattices(esk1_0,X1,esk3_0)=X1|~r1_lattices(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_15]), c_0_16])]), c_0_17])).
cnf(c_0_24, negated_conjecture, (r1_lattices(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_25, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_26, plain, ![X27, X28, X29]:((~v9_lattices(X27)|(~m1_subset_1(X28,u1_struct_0(X27))|(~m1_subset_1(X29,u1_struct_0(X27))|k2_lattices(X27,X28,k1_lattices(X27,X28,X29))=X28))|(v2_struct_0(X27)|~l3_lattices(X27)))&((m1_subset_1(esk10_1(X27),u1_struct_0(X27))|v9_lattices(X27)|(v2_struct_0(X27)|~l3_lattices(X27)))&((m1_subset_1(esk11_1(X27),u1_struct_0(X27))|v9_lattices(X27)|(v2_struct_0(X27)|~l3_lattices(X27)))&(k2_lattices(X27,esk10_1(X27),k1_lattices(X27,esk10_1(X27),esk11_1(X27)))!=esk10_1(X27)|v9_lattices(X27)|(v2_struct_0(X27)|~l3_lattices(X27)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d9_lattices])])])])])])).
cnf(c_0_27, plain, (k1_lattices(X1,k2_lattices(X1,X2,X3),X3)=X3|v2_struct_0(X1)|~v8_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, negated_conjecture, (m1_subset_1(k2_lattices(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))|~l1_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_17])).
cnf(c_0_29, negated_conjecture, (k2_lattices(esk1_0,k2_lattices(esk1_0,X1,X2),esk4_0)=k2_lattices(esk1_0,X1,k2_lattices(esk1_0,X2,esk4_0))|~l1_lattices(esk1_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_20]), c_0_22])]), c_0_17])).
cnf(c_0_30, negated_conjecture, (k2_lattices(esk1_0,esk2_0,esk3_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])])).
cnf(c_0_31, plain, (k2_lattices(X1,X2,k1_lattices(X1,X2,X3))=X2|v2_struct_0(X1)|~v9_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, negated_conjecture, (k1_lattices(esk1_0,k2_lattices(esk1_0,X1,k2_lattices(esk1_0,X2,esk4_0)),k2_lattices(esk1_0,X2,esk4_0))=k2_lattices(esk1_0,X2,esk4_0)|~l1_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_14]), c_0_16])]), c_0_17])).
cnf(c_0_33, negated_conjecture, (k2_lattices(esk1_0,esk2_0,k2_lattices(esk1_0,esk3_0,esk4_0))=k2_lattices(esk1_0,esk2_0,esk4_0)|~l1_lattices(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_13]), c_0_25])])).
cnf(c_0_34, negated_conjecture, (m1_subset_1(k2_lattices(esk1_0,X1,k2_lattices(esk1_0,X2,esk4_0)),u1_struct_0(esk1_0))|~l1_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_28]), c_0_17])).
cnf(c_0_35, negated_conjecture, (k2_lattices(esk1_0,X1,k1_lattices(esk1_0,X1,k2_lattices(esk1_0,X2,esk4_0)))=X1|~l1_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_28]), c_0_14]), c_0_15])]), c_0_17])).
cnf(c_0_36, negated_conjecture, (k1_lattices(esk1_0,k2_lattices(esk1_0,esk2_0,esk4_0),k2_lattices(esk1_0,esk3_0,esk4_0))=k2_lattices(esk1_0,esk3_0,esk4_0)|~l1_lattices(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_25]), c_0_13])])).
cnf(c_0_37, negated_conjecture, (m1_subset_1(k2_lattices(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))|~l1_lattices(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_33]), c_0_25]), c_0_13])])).
cnf(c_0_38, plain, (r1_lattices(X1,X2,X3)|v2_struct_0(X1)|k2_lattices(X1,X2,X3)!=X2|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_39, negated_conjecture, (k2_lattices(esk1_0,k2_lattices(esk1_0,esk2_0,esk4_0),k2_lattices(esk1_0,esk3_0,esk4_0))=k2_lattices(esk1_0,esk2_0,esk4_0)|~l1_lattices(esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_13])]), c_0_37])).
cnf(c_0_40, negated_conjecture, (~r1_lattices(esk1_0,k2_lattices(esk1_0,esk2_0,esk4_0),k2_lattices(esk1_0,esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_41, negated_conjecture, (~l1_lattices(esk1_0)|~m1_subset_1(k2_lattices(esk1_0,esk3_0,esk4_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_14]), c_0_15]), c_0_16])]), c_0_40]), c_0_17]), c_0_37])).
fof(c_0_42, plain, ![X32]:((l1_lattices(X32)|~l3_lattices(X32))&(l2_lattices(X32)|~l3_lattices(X32))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
cnf(c_0_43, negated_conjecture, (~l1_lattices(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_28]), c_0_13])])).
cnf(c_0_44, plain, (l1_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_45, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_14])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 46
# Proof object clause steps            : 31
# Proof object formula steps           : 15
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 43
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 28
# Processed clauses                    : 421
# ...of these trivial                  : 5
# ...subsumed                          : 176
# ...remaining for further processing  : 240
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 42
# Backward-rewritten                   : 0
# Generated clauses                    : 2024
# ...of the previous two non-trivial   : 1165
# Contextual simplify-reflections      : 27
# Paramodulations                      : 2024
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 170
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 157
# Current number of unprocessed clauses: 799
# ...number of literals in the above   : 4173
# Current number of archived formulas  : 0
# Current number of archived clauses   : 70
# Clause-clause subsumption calls (NU) : 14247
# Rec. Clause-clause subsumption calls : 4721
# Non-unit clause-clause subsumptions  : 245
# Unit Clause-clause subsumption calls : 160
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 65773

# -------------------------------------------------
# User time                : 0.027 s
# System time              : 0.007 s
# Total time               : 0.033 s
# Maximum resident set size: 28668 pages
