# No SInE strategy applied
# Auto-Mode selected heuristic G_E___200_B02_F1_SE_CS_SP_PI_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(cc1_relat_1, axiom, ![X1]:(v1_xboole_0(X1)=>v1_relat_1(X1)), file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', cc1_relat_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', fc1_xboole_0)).
fof(d2_xboole_0, axiom, k1_xboole_0=o_0_0_xboole_0, file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', d2_xboole_0)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', t1_boole)).
fof(t13_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>k1_relat_1(k2_xboole_0(X1,X2))=k2_xboole_0(k1_relat_1(X1),k1_relat_1(X2)))), file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', t13_relat_1)).
fof(t24_relat_1, axiom, ![X1, X2, X3, X4, X5]:(v1_relat_1(X5)=>(X5=k2_tarski(k4_tarski(X1,X2),k4_tarski(X3,X4))=>(k1_relat_1(X5)=k2_tarski(X1,X3)&k2_relat_1(X5)=k2_tarski(X2,X4)))), file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', t24_relat_1)).
fof(fc7_relat_1, axiom, ![X1, X2, X3, X4]:v1_relat_1(k2_tarski(k4_tarski(X1,X2),k4_tarski(X3,X4))), file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', fc7_relat_1)).
fof(t64_zfmisc_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k4_xboole_0(X2,k1_tarski(X3)))<=>(r2_hidden(X1,X2)&~(X1=X3))), file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', t64_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', t69_enumset1)).
fof(t49_zfmisc_1, axiom, ![X1, X2]:~(k2_xboole_0(k1_tarski(X1),X2)=k1_xboole_0), file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', t49_zfmisc_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', t7_xboole_0)).
fof(t46_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0, file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', t46_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', commutativity_k2_xboole_0)).
fof(t4_boole, axiom, ![X1]:k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0, file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', t4_boole)).
fof(t60_relat_1, conjecture, (k1_relat_1(k1_xboole_0)=k1_xboole_0&k2_relat_1(k1_xboole_0)=k1_xboole_0), file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', t60_relat_1)).
fof(t37_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>(k2_relat_1(X1)=k1_relat_1(k4_relat_1(X1))&k1_relat_1(X1)=k2_relat_1(k4_relat_1(X1)))), file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', t37_relat_1)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', t2_boole)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', t12_setfam_1)).
fof(t22_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k3_xboole_0(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))=X1), file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', t22_relat_1)).
fof(dt_k4_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>v1_relat_1(k4_relat_1(X1))), file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', dt_k4_relat_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/0d8e70570d154f9e97f632c7d23df664.E_input', t113_zfmisc_1)).
fof(c_0_21, plain, ![X6]:(~v1_xboole_0(X6)|v1_relat_1(X6)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relat_1])])).
cnf(c_0_22, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_23, plain, (k1_xboole_0=o_0_0_xboole_0), inference(split_conjunct,[status(thm)],[d2_xboole_0])).
fof(c_0_24, plain, ![X7]:k2_xboole_0(X7,k1_xboole_0)=X7, inference(variable_rename,[status(thm)],[t1_boole])).
fof(c_0_25, plain, ![X8, X9]:(~v1_relat_1(X8)|(~v1_relat_1(X9)|k1_relat_1(k2_xboole_0(X8,X9))=k2_xboole_0(k1_relat_1(X8),k1_relat_1(X9)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_relat_1])])])).
cnf(c_0_26, plain, (v1_relat_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (v1_xboole_0(o_0_0_xboole_0)), inference(rw,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_28, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_29, plain, ![X10, X11, X12, X13, X14]:((k1_relat_1(X14)=k2_tarski(X10,X12)|X14!=k2_tarski(k4_tarski(X10,X11),k4_tarski(X12,X13))|~v1_relat_1(X14))&(k2_relat_1(X14)=k2_tarski(X11,X13)|X14!=k2_tarski(k4_tarski(X10,X11),k4_tarski(X12,X13))|~v1_relat_1(X14))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t24_relat_1])])])).
fof(c_0_30, plain, ![X15, X16, X17, X18]:v1_relat_1(k2_tarski(k4_tarski(X15,X16),k4_tarski(X17,X18))), inference(variable_rename,[status(thm)],[fc7_relat_1])).
fof(c_0_31, plain, ![X19, X20, X21]:(((r2_hidden(X19,X20)|~r2_hidden(X19,k4_xboole_0(X20,k1_tarski(X21))))&(X19!=X21|~r2_hidden(X19,k4_xboole_0(X20,k1_tarski(X21)))))&(~r2_hidden(X19,X20)|X19=X21|r2_hidden(X19,k4_xboole_0(X20,k1_tarski(X21))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t64_zfmisc_1])])])])).
fof(c_0_32, plain, ![X22]:k2_tarski(X22,X22)=k1_tarski(X22), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_33, plain, ![X23, X24]:k2_xboole_0(k1_tarski(X23),X24)!=k1_xboole_0, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t49_zfmisc_1])])).
fof(c_0_34, plain, ![X25]:(X25=k1_xboole_0|r2_hidden(esk1_1(X25),X25)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
fof(c_0_35, plain, ![X27, X28]:k4_xboole_0(X27,k2_xboole_0(X27,X28))=k1_xboole_0, inference(variable_rename,[status(thm)],[t46_xboole_1])).
cnf(c_0_36, plain, (k1_relat_1(k2_xboole_0(X1,X2))=k2_xboole_0(k1_relat_1(X1),k1_relat_1(X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_37, plain, (v1_relat_1(o_0_0_xboole_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_38, plain, (k2_xboole_0(X1,o_0_0_xboole_0)=X1), inference(rw,[status(thm)],[c_0_28, c_0_23])).
cnf(c_0_39, plain, (k1_relat_1(X1)=k2_tarski(X2,X3)|X1!=k2_tarski(k4_tarski(X2,X4),k4_tarski(X3,X5))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_40, plain, (v1_relat_1(k2_tarski(k4_tarski(X1,X2),k4_tarski(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_41, plain, ![X29, X30]:k2_xboole_0(X29,X30)=k2_xboole_0(X30,X29), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_42, plain, (X1!=X2|~r2_hidden(X1,k4_xboole_0(X3,k1_tarski(X2)))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_43, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_44, plain, ![X31]:k4_xboole_0(k1_xboole_0,X31)=k1_xboole_0, inference(variable_rename,[status(thm)],[t4_boole])).
fof(c_0_45, negated_conjecture, ~((k1_relat_1(k1_xboole_0)=k1_xboole_0&k2_relat_1(k1_xboole_0)=k1_xboole_0)), inference(assume_negation,[status(cth)],[t60_relat_1])).
cnf(c_0_46, plain, (k2_xboole_0(k1_tarski(X1),X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_47, plain, (X1=X3|r2_hidden(X1,k4_xboole_0(X2,k1_tarski(X3)))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_48, plain, (X1=k1_xboole_0|r2_hidden(esk1_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_49, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_50, plain, (k2_xboole_0(k1_relat_1(X1),k1_relat_1(o_0_0_xboole_0))=k1_relat_1(X1)|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_51, plain, (k1_relat_1(k2_tarski(k4_tarski(X1,X2),k4_tarski(X3,X4)))=k2_tarski(X1,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_39]), c_0_40])])).
cnf(c_0_52, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_53, plain, (X1!=X2|~r2_hidden(X1,k4_xboole_0(X3,k2_tarski(X2,X2)))), inference(rw,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_54, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_55, negated_conjecture, (k1_relat_1(k1_xboole_0)!=k1_xboole_0|k2_relat_1(k1_xboole_0)!=k1_xboole_0), inference(fof_nnf,[status(thm)],[c_0_45])).
fof(c_0_56, plain, ![X32]:((k2_relat_1(X32)=k1_relat_1(k4_relat_1(X32))|~v1_relat_1(X32))&(k1_relat_1(X32)=k2_relat_1(k4_relat_1(X32))|~v1_relat_1(X32))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t37_relat_1])])])).
cnf(c_0_57, plain, (k2_xboole_0(k2_tarski(X1,X1),X2)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_46, c_0_43])).
cnf(c_0_58, plain, (X1=X3|r2_hidden(X1,k4_xboole_0(X2,k2_tarski(X3,X3)))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_47, c_0_43])).
cnf(c_0_59, plain, (X1=o_0_0_xboole_0|r2_hidden(esk1_1(X1),X1)), inference(rw,[status(thm)],[c_0_48, c_0_23])).
cnf(c_0_60, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=o_0_0_xboole_0), inference(rw,[status(thm)],[c_0_49, c_0_23])).
cnf(c_0_61, plain, (k2_xboole_0(k1_relat_1(o_0_0_xboole_0),k2_tarski(X1,X2))=k2_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_40]), c_0_51]), c_0_51]), c_0_52])).
cnf(c_0_62, plain, (~r2_hidden(X1,k4_xboole_0(X2,k2_tarski(X1,X1)))), inference(er,[status(thm)],[c_0_53])).
cnf(c_0_63, plain, (k4_xboole_0(o_0_0_xboole_0,X1)=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_23]), c_0_23])).
fof(c_0_64, plain, ![X33]:k3_xboole_0(X33,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_65, plain, ![X34, X35]:k1_setfam_1(k2_tarski(X34,X35))=k3_xboole_0(X34,X35), inference(variable_rename,[status(thm)],[t12_setfam_1])).
cnf(c_0_66, negated_conjecture, (k1_relat_1(k1_xboole_0)!=k1_xboole_0|k2_relat_1(k1_xboole_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_67, plain, (k2_relat_1(X1)=k1_relat_1(k4_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_68, plain, (k2_xboole_0(k2_tarski(X1,X1),X2)!=o_0_0_xboole_0), inference(rw,[status(thm)],[c_0_57, c_0_23])).
cnf(c_0_69, plain, (esk1_1(X1)=X2|X1=o_0_0_xboole_0|r2_hidden(esk1_1(X1),k4_xboole_0(X1,k2_tarski(X2,X2)))), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_70, plain, (k4_xboole_0(k1_relat_1(o_0_0_xboole_0),k2_tarski(X1,X2))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_71, plain, (~r2_hidden(X1,o_0_0_xboole_0)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
fof(c_0_72, plain, ![X36]:(~v1_relat_1(X36)|k3_xboole_0(X36,k2_zfmisc_1(k1_relat_1(X36),k2_relat_1(X36)))=X36), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t22_relat_1])])).
fof(c_0_73, plain, ![X37]:(~v1_relat_1(X37)|v1_relat_1(k4_relat_1(X37))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_relat_1])])).
fof(c_0_74, plain, ![X38, X39]:((k2_zfmisc_1(X38,X39)!=k1_xboole_0|(X38=k1_xboole_0|X39=k1_xboole_0))&((X38!=k1_xboole_0|k2_zfmisc_1(X38,X39)=k1_xboole_0)&(X39!=k1_xboole_0|k2_zfmisc_1(X38,X39)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_75, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_76, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_77, negated_conjecture, (k1_relat_1(o_0_0_xboole_0)!=o_0_0_xboole_0|k2_relat_1(o_0_0_xboole_0)!=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_23]), c_0_23]), c_0_23]), c_0_23])).
cnf(c_0_78, plain, (k2_relat_1(o_0_0_xboole_0)=k1_relat_1(k4_relat_1(o_0_0_xboole_0))), inference(spm,[status(thm)],[c_0_67, c_0_37])).
cnf(c_0_79, plain, (k2_tarski(X1,X1)!=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_68, c_0_38])).
cnf(c_0_80, plain, (esk1_1(k1_relat_1(o_0_0_xboole_0))=X1|k1_relat_1(o_0_0_xboole_0)=o_0_0_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_71])).
cnf(c_0_81, plain, (k3_xboole_0(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))=X1|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_82, plain, (v1_relat_1(k4_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_83, plain, (k1_relat_1(X1)=k2_relat_1(k4_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_84, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_85, plain, (k1_setfam_1(k2_tarski(X1,k1_xboole_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_86, negated_conjecture, (k1_relat_1(k4_relat_1(o_0_0_xboole_0))!=o_0_0_xboole_0|k1_relat_1(o_0_0_xboole_0)!=o_0_0_xboole_0), inference(rw,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_87, plain, (k1_relat_1(o_0_0_xboole_0)=o_0_0_xboole_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_80])).
cnf(c_0_88, plain, (k1_setfam_1(k2_tarski(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))=X1|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_81, c_0_76])).
cnf(c_0_89, plain, (v1_relat_1(k4_relat_1(o_0_0_xboole_0))), inference(spm,[status(thm)],[c_0_82, c_0_37])).
cnf(c_0_90, plain, (k2_relat_1(k4_relat_1(o_0_0_xboole_0))=k1_relat_1(o_0_0_xboole_0)), inference(spm,[status(thm)],[c_0_83, c_0_37])).
cnf(c_0_91, plain, (k2_zfmisc_1(X1,o_0_0_xboole_0)=o_0_0_xboole_0), inference(er,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_84, c_0_23]), c_0_23])])).
cnf(c_0_92, plain, (k1_setfam_1(k2_tarski(X1,o_0_0_xboole_0))=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_23]), c_0_23])).
cnf(c_0_93, negated_conjecture, (k1_relat_1(k4_relat_1(o_0_0_xboole_0))!=o_0_0_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_87])])).
cnf(c_0_94, plain, (k4_relat_1(o_0_0_xboole_0)=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_90]), c_0_87]), c_0_91]), c_0_92])).
cnf(c_0_95, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_94]), c_0_87])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 96
# Proof object clause steps            : 55
# Proof object formula steps           : 41
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 21
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 41
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 25
# Processed clauses                    : 129
# ...of these trivial                  : 9
# ...subsumed                          : 11
# ...remaining for further processing  : 109
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 59
# Generated clauses                    : 385
# ...of the previous two non-trivial   : 358
# Contextual simplify-reflections      : 1
# Paramodulations                      : 380
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 45
#    Positive orientable unit clauses  : 25
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 14
# Current number of unprocessed clauses: 138
# ...number of literals in the above   : 162
# Current number of archived formulas  : 0
# Current number of archived clauses   : 61
# Clause-clause subsumption calls (NU) : 78
# Rec. Clause-clause subsumption calls : 70
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 81
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 154
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6988

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28560 pages
