# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/0e48af1922df4e6186c4ce4d13074158.E_input', d5_xboole_0)).
fof(t4_boole, axiom, ![X1]:k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0, file('data/MPTP/proofs/0e48af1922df4e6186c4ce4d13074158.E_input', t4_boole)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/0e48af1922df4e6186c4ce4d13074158.E_input', l32_xboole_1)).
fof(t73_zfmisc_1, conjecture, ![X1, X2, X3]:(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_xboole_0<=>(r2_hidden(X1,X3)&r2_hidden(X2,X3))), file('data/MPTP/proofs/0e48af1922df4e6186c4ce4d13074158.E_input', t73_zfmisc_1)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/0e48af1922df4e6186c4ce4d13074158.E_input', d2_tarski)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/0e48af1922df4e6186c4ce4d13074158.E_input', t70_enumset1)).
fof(c_0_6, plain, ![X21, X22, X23, X24, X25, X26, X27, X28]:((((r2_hidden(X24,X21)|~r2_hidden(X24,X23)|X23!=k4_xboole_0(X21,X22))&(~r2_hidden(X24,X22)|~r2_hidden(X24,X23)|X23!=k4_xboole_0(X21,X22)))&(~r2_hidden(X25,X21)|r2_hidden(X25,X22)|r2_hidden(X25,X23)|X23!=k4_xboole_0(X21,X22)))&((~r2_hidden(esk5_3(X26,X27,X28),X28)|(~r2_hidden(esk5_3(X26,X27,X28),X26)|r2_hidden(esk5_3(X26,X27,X28),X27))|X28=k4_xboole_0(X26,X27))&((r2_hidden(esk5_3(X26,X27,X28),X26)|r2_hidden(esk5_3(X26,X27,X28),X28)|X28=k4_xboole_0(X26,X27))&(~r2_hidden(esk5_3(X26,X27,X28),X27)|r2_hidden(esk5_3(X26,X27,X28),X28)|X28=k4_xboole_0(X26,X27))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_7, plain, ![X30]:k4_xboole_0(k1_xboole_0,X30)=k1_xboole_0, inference(variable_rename,[status(thm)],[t4_boole])).
fof(c_0_8, plain, ![X19, X20]:((k4_xboole_0(X19,X20)!=k1_xboole_0|r1_tarski(X19,X20))&(~r1_tarski(X19,X20)|k4_xboole_0(X19,X20)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_9, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_11, plain, (~epred4_0<=>![X4, X2]:~r1_tarski(X2,X4)), introduced(definition)).
cnf(c_0_12, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_13, plain, (~epred3_0<=>![X1, X3]:(~r2_hidden(X1,X3)|X3!=k1_xboole_0)), introduced(definition)).
cnf(c_0_14, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)|~r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_17, plain, (epred4_0|~r1_tarski(X1,X2)), inference(split_equiv,[status(thm)],[c_0_11])).
cnf(c_0_18, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_12, c_0_10])).
fof(c_0_19, negated_conjecture, ~(![X1, X2, X3]:(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_xboole_0<=>(r2_hidden(X1,X3)&r2_hidden(X2,X3)))), inference(assume_negation,[status(cth)],[t73_zfmisc_1])).
fof(c_0_20, plain, ![X8, X9, X10, X11, X12, X13, X14, X15]:(((~r2_hidden(X11,X10)|(X11=X8|X11=X9)|X10!=k2_tarski(X8,X9))&((X12!=X8|r2_hidden(X12,X10)|X10!=k2_tarski(X8,X9))&(X12!=X9|r2_hidden(X12,X10)|X10!=k2_tarski(X8,X9))))&(((esk4_3(X13,X14,X15)!=X13|~r2_hidden(esk4_3(X13,X14,X15),X15)|X15=k2_tarski(X13,X14))&(esk4_3(X13,X14,X15)!=X14|~r2_hidden(esk4_3(X13,X14,X15),X15)|X15=k2_tarski(X13,X14)))&(r2_hidden(esk4_3(X13,X14,X15),X15)|(esk4_3(X13,X14,X15)=X13|esk4_3(X13,X14,X15)=X14)|X15=k2_tarski(X13,X14)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_21, plain, ![X17, X18]:k1_enumset1(X17,X17,X18)=k2_tarski(X17,X18), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_22, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_23, plain, (~epred4_0|~epred3_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_13]), c_0_11])).
cnf(c_0_24, plain, (epred4_0), inference(spm,[status(thm)],[c_0_17, c_0_18])).
fof(c_0_25, negated_conjecture, ((k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)!=k1_xboole_0|(~r2_hidden(esk1_0,esk3_0)|~r2_hidden(esk2_0,esk3_0)))&((r2_hidden(esk1_0,esk3_0)|k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)=k1_xboole_0)&(r2_hidden(esk2_0,esk3_0)|k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])).
cnf(c_0_26, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (epred3_0|X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(split_equiv,[status(thm)],[c_0_13])).
cnf(c_0_29, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_30, plain, (~epred3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_24])])).
cnf(c_0_31, negated_conjecture, (r2_hidden(esk1_0,esk3_0)|k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X2,X2,X4)), inference(rw,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_33, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_34, plain, (r2_hidden(X1,X2)|k4_xboole_0(X3,X2)!=k1_xboole_0|~r2_hidden(X1,X3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])).
cnf(c_0_35, negated_conjecture, (k4_xboole_0(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0)=k1_xboole_0|r2_hidden(esk1_0,esk3_0)), inference(rw,[status(thm)],[c_0_31, c_0_27])).
cnf(c_0_36, plain, (r2_hidden(X1,X2)|X2!=k1_enumset1(X1,X1,X3)), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_37, negated_conjecture, (r2_hidden(esk2_0,esk3_0)|k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_38, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X4,X2)), inference(rw,[status(thm)],[c_0_33, c_0_27])).
cnf(c_0_39, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_9])).
cnf(c_0_40, negated_conjecture, (k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)!=k1_xboole_0|~r2_hidden(esk1_0,esk3_0)|~r2_hidden(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_41, negated_conjecture, (r2_hidden(esk1_0,esk3_0)|r2_hidden(X1,esk3_0)|~r2_hidden(X1,k1_enumset1(esk1_0,esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_42, plain, (r2_hidden(X1,k1_enumset1(X1,X1,X2))), inference(er,[status(thm)],[c_0_36])).
cnf(c_0_43, negated_conjecture, (k4_xboole_0(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0)=k1_xboole_0|r2_hidden(esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_37, c_0_27])).
cnf(c_0_44, plain, (r2_hidden(X1,X2)|X2!=k1_enumset1(X3,X3,X1)), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_45, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_46, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_39, c_0_10])).
cnf(c_0_47, negated_conjecture, (k4_xboole_0(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0)!=k1_xboole_0|~r2_hidden(esk1_0,esk3_0)|~r2_hidden(esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_40, c_0_27])).
cnf(c_0_48, negated_conjecture, (r2_hidden(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_49, negated_conjecture, (r2_hidden(esk2_0,esk3_0)|r2_hidden(X1,esk3_0)|~r2_hidden(X1,k1_enumset1(esk1_0,esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_34, c_0_43])).
cnf(c_0_50, plain, (r2_hidden(X1,k1_enumset1(X2,X2,X1))), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_51, plain, (X1=X4|X1=X3|X2!=k1_enumset1(X3,X3,X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_45, c_0_27])).
cnf(c_0_52, plain, (~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_10]), c_0_46])).
cnf(c_0_53, plain, (r2_hidden(esk5_3(X1,X2,X3),X1)|r2_hidden(esk5_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_54, negated_conjecture, (k4_xboole_0(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0)!=k1_xboole_0|~r2_hidden(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_48])])).
cnf(c_0_55, negated_conjecture, (r2_hidden(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_56, plain, (X1=X2|X1=X3|~r2_hidden(X1,k1_enumset1(X2,X2,X3))), inference(er,[status(thm)],[c_0_51])).
cnf(c_0_57, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|r2_hidden(esk5_3(X1,X2,k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_58, negated_conjecture, (k4_xboole_0(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_55])])).
cnf(c_0_59, plain, (esk5_3(k1_enumset1(X1,X1,X2),X3,k1_xboole_0)=X2|esk5_3(k1_enumset1(X1,X1,X2),X3,k1_xboole_0)=X1|k4_xboole_0(k1_enumset1(X1,X1,X2),X3)=k1_xboole_0), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_60, plain, (r2_hidden(esk5_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk5_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_61, negated_conjecture, (esk5_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_xboole_0)=esk1_0|esk5_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_xboole_0)=esk2_0), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_62, negated_conjecture, (esk5_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_xboole_0)=esk1_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_55])]), c_0_58]), c_0_52])).
cnf(c_0_63, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_62]), c_0_48])]), c_0_58]), c_0_52]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 64
# Proof object clause steps            : 49
# Proof object formula steps           : 15
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 27
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 18
# Processed clauses                    : 1240
# ...of these trivial                  : 16
# ...subsumed                          : 945
# ...remaining for further processing  : 279
# Other redundant clauses eliminated   : 456
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 19
# Backward-rewritten                   : 31
# Generated clauses                    : 11835
# ...of the previous two non-trivial   : 10618
# Contextual simplify-reflections      : 17
# Paramodulations                      : 11188
# Factorizations                       : 121
# Equation resolutions                 : 500
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 199
#    Positive orientable unit clauses  : 26
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 12
#    Non-unit-clauses                  : 161
# Current number of unprocessed clauses: 9325
# ...number of literals in the above   : 101947
# Current number of archived formulas  : 0
# Current number of archived clauses   : 69
# Clause-clause subsumption calls (NU) : 21847
# Rec. Clause-clause subsumption calls : 8665
# Non-unit clause-clause subsumptions  : 852
# Unit Clause-clause subsumption calls : 809
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 77
# BW rewrite match successes           : 13
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 270004

# -------------------------------------------------
# User time                : 0.355 s
# System time              : 0.010 s
# Total time               : 0.366 s
# Maximum resident set size: 28548 pages
