# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t49_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v15_waybel_0(X2,X1)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&X2=k6_waybel_0(X1,X3))))), file('data/MPTP/proofs/0eb472e81b1d4636a17cbd81f3b13d34.E_input', t49_waybel_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/0eb472e81b1d4636a17cbd81f3b13d34.E_input', t4_subset)).
fof(dt_k6_waybel_0, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k6_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/0eb472e81b1d4636a17cbd81f3b13d34.E_input', dt_k6_waybel_0)).
fof(reflexivity_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>r3_orders_2(X1,X2,X2)), file('data/MPTP/proofs/0eb472e81b1d4636a17cbd81f3b13d34.E_input', reflexivity_r3_orders_2)).
fof(t18_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k6_waybel_0(X1,X2))<=>r1_orders_2(X1,X2,X3))))), file('data/MPTP/proofs/0eb472e81b1d4636a17cbd81f3b13d34.E_input', t18_waybel_0)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/0eb472e81b1d4636a17cbd81f3b13d34.E_input', redefinition_r3_orders_2)).
fof(d20_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v13_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r1_orders_2(X1,X3,X4))=>r2_hidden(X4,X2))))))), file('data/MPTP/proofs/0eb472e81b1d4636a17cbd81f3b13d34.E_input', d20_waybel_0)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/0eb472e81b1d4636a17cbd81f3b13d34.E_input', d8_lattice3)).
fof(d22_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v15_waybel_0(X2,X1)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_hidden(X3,X2)&r1_lattice3(X1,X2,X3)))))), file('data/MPTP/proofs/0eb472e81b1d4636a17cbd81f3b13d34.E_input', d22_waybel_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/0eb472e81b1d4636a17cbd81f3b13d34.E_input', t24_orders_2)).
fof(t7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,X2)=>r2_hidden(X4,X3)))=>r1_tarski(X2,X3)))), file('data/MPTP/proofs/0eb472e81b1d4636a17cbd81f3b13d34.E_input', t7_subset_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/0eb472e81b1d4636a17cbd81f3b13d34.E_input', d10_xboole_0)).
fof(c_0_12, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v15_waybel_0(X2,X1)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&X2=k6_waybel_0(X1,X3)))))), inference(assume_negation,[status(cth)],[t49_waybel_0])).
fof(c_0_13, plain, ![X9, X10, X11]:(~r2_hidden(X9,X10)|~m1_subset_1(X10,k1_zfmisc_1(X11))|m1_subset_1(X9,X11)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_14, plain, ![X12, X13]:(v2_struct_0(X12)|~l1_orders_2(X12)|~m1_subset_1(X13,u1_struct_0(X12))|m1_subset_1(k6_waybel_0(X12,X13),k1_zfmisc_1(u1_struct_0(X12)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_waybel_0])])])).
fof(c_0_15, plain, ![X14, X15, X16]:(v2_struct_0(X14)|(~v3_orders_2(X14)|~l1_orders_2(X14))|(~m1_subset_1(X15,u1_struct_0(X14))|~m1_subset_1(X16,u1_struct_0(X14)))|r3_orders_2(X14,X15,X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r3_orders_2])])])).
fof(c_0_16, negated_conjecture, ![X7]:((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0))))&((~v1_xboole_0(esk2_0)&(v2_waybel_0(esk2_0,esk1_0)&(v13_waybel_0(esk2_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))))&((~v15_waybel_0(esk2_0,esk1_0)|(~m1_subset_1(X7,u1_struct_0(esk1_0))|esk2_0!=k6_waybel_0(esk1_0,X7)))&((m1_subset_1(esk3_0,u1_struct_0(esk1_0))|v15_waybel_0(esk2_0,esk1_0))&(esk2_0=k6_waybel_0(esk1_0,esk3_0)|v15_waybel_0(esk2_0,esk1_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])])])).
fof(c_0_17, plain, ![X17, X18, X19]:((~r2_hidden(X19,k6_waybel_0(X17,X18))|r1_orders_2(X17,X18,X19)|~m1_subset_1(X19,u1_struct_0(X17))|~m1_subset_1(X18,u1_struct_0(X17))|(v2_struct_0(X17)|~l1_orders_2(X17)))&(~r1_orders_2(X17,X18,X19)|r2_hidden(X19,k6_waybel_0(X17,X18))|~m1_subset_1(X19,u1_struct_0(X17))|~m1_subset_1(X18,u1_struct_0(X17))|(v2_struct_0(X17)|~l1_orders_2(X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t18_waybel_0])])])])])).
cnf(c_0_18, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (v2_struct_0(X1)|m1_subset_1(k6_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_20, plain, ![X20, X21, X22]:((~r3_orders_2(X20,X21,X22)|r1_orders_2(X20,X21,X22)|(v2_struct_0(X20)|(~v3_orders_2(X20)|~l1_orders_2(X20))|(~m1_subset_1(X21,u1_struct_0(X20))|~m1_subset_1(X22,u1_struct_0(X20)))))&(~r1_orders_2(X20,X21,X22)|r3_orders_2(X20,X21,X22)|(v2_struct_0(X20)|(~v3_orders_2(X20)|~l1_orders_2(X20))|(~m1_subset_1(X21,u1_struct_0(X20))|~m1_subset_1(X22,u1_struct_0(X20)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_21, plain, (v2_struct_0(X1)|r3_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))|v15_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_26, plain, ![X23, X24, X25, X26]:((~v13_waybel_0(X24,X23)|(~m1_subset_1(X25,u1_struct_0(X23))|(~m1_subset_1(X26,u1_struct_0(X23))|(~r2_hidden(X25,X24)|~r1_orders_2(X23,X25,X26)|r2_hidden(X26,X24))))|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X23)))|~l1_orders_2(X23))&((m1_subset_1(esk4_2(X23,X24),u1_struct_0(X23))|v13_waybel_0(X24,X23)|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X23)))|~l1_orders_2(X23))&((m1_subset_1(esk5_2(X23,X24),u1_struct_0(X23))|v13_waybel_0(X24,X23)|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X23)))|~l1_orders_2(X23))&(((r2_hidden(esk4_2(X23,X24),X24)|v13_waybel_0(X24,X23)|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X23)))|~l1_orders_2(X23))&(r1_orders_2(X23,esk4_2(X23,X24),esk5_2(X23,X24))|v13_waybel_0(X24,X23)|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X23)))|~l1_orders_2(X23)))&(~r2_hidden(esk5_2(X23,X24),X24)|v13_waybel_0(X24,X23)|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X23)))|~l1_orders_2(X23)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d20_waybel_0])])])])])).
cnf(c_0_27, plain, (r1_orders_2(X2,X3,X1)|v2_struct_0(X2)|~r2_hidden(X1,k6_waybel_0(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k6_waybel_0(X2,X3))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
fof(c_0_29, plain, ![X29, X30, X31, X32]:((~r1_lattice3(X29,X30,X31)|(~m1_subset_1(X32,u1_struct_0(X29))|(~r2_hidden(X32,X30)|r1_orders_2(X29,X31,X32)))|~m1_subset_1(X31,u1_struct_0(X29))|~l1_orders_2(X29))&((m1_subset_1(esk6_3(X29,X30,X31),u1_struct_0(X29))|r1_lattice3(X29,X30,X31)|~m1_subset_1(X31,u1_struct_0(X29))|~l1_orders_2(X29))&((r2_hidden(esk6_3(X29,X30,X31),X30)|r1_lattice3(X29,X30,X31)|~m1_subset_1(X31,u1_struct_0(X29))|~l1_orders_2(X29))&(~r1_orders_2(X29,X31,esk6_3(X29,X30,X31))|r1_lattice3(X29,X30,X31)|~m1_subset_1(X31,u1_struct_0(X29))|~l1_orders_2(X29))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
cnf(c_0_30, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_31, negated_conjecture, (r3_orders_2(esk1_0,X1,X1)|v15_waybel_0(esk2_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24])]), c_0_25])).
fof(c_0_32, plain, ![X34, X35, X37]:(((m1_subset_1(esk7_2(X34,X35),u1_struct_0(X34))|~v15_waybel_0(X35,X34)|(v1_xboole_0(X35)|(~v2_waybel_0(X35,X34)|(~v13_waybel_0(X35,X34)|~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X34))))))|(v2_struct_0(X34)|(~v3_orders_2(X34)|(~v4_orders_2(X34)|~l1_orders_2(X34)))))&((r2_hidden(esk7_2(X34,X35),X35)|~v15_waybel_0(X35,X34)|(v1_xboole_0(X35)|(~v2_waybel_0(X35,X34)|(~v13_waybel_0(X35,X34)|~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X34))))))|(v2_struct_0(X34)|(~v3_orders_2(X34)|(~v4_orders_2(X34)|~l1_orders_2(X34)))))&(r1_lattice3(X34,X35,esk7_2(X34,X35))|~v15_waybel_0(X35,X34)|(v1_xboole_0(X35)|(~v2_waybel_0(X35,X34)|(~v13_waybel_0(X35,X34)|~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X34))))))|(v2_struct_0(X34)|(~v3_orders_2(X34)|(~v4_orders_2(X34)|~l1_orders_2(X34)))))))&(~m1_subset_1(X37,u1_struct_0(X34))|(~r2_hidden(X37,X35)|~r1_lattice3(X34,X35,X37))|v15_waybel_0(X35,X34)|(v1_xboole_0(X35)|(~v2_waybel_0(X35,X34)|(~v13_waybel_0(X35,X34)|~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X34))))))|(v2_struct_0(X34)|(~v3_orders_2(X34)|(~v4_orders_2(X34)|~l1_orders_2(X34)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d22_waybel_0])])])])])])).
cnf(c_0_33, plain, (r2_hidden(X4,X1)|~v13_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,X1)|~r1_orders_2(X2,X3,X4)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r2_hidden(X3,k6_waybel_0(X1,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_35, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, negated_conjecture, (esk2_0=k6_waybel_0(esk1_0,esk3_0)|v15_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_37, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|v15_waybel_0(esk2_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_23]), c_0_24])]), c_0_25])).
cnf(c_0_38, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, plain, (r1_lattice3(X1,X2,esk7_2(X1,X2))|v1_xboole_0(X2)|v2_struct_0(X1)|~v15_waybel_0(X2,X1)|~v2_waybel_0(X2,X1)|~v13_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40, plain, (m1_subset_1(esk7_2(X1,X2),u1_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~v15_waybel_0(X2,X1)|~v2_waybel_0(X2,X1)|~v13_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_41, plain, (r2_hidden(X1,X2)|~r1_orders_2(X3,X4,X1)|~r2_hidden(X4,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~v13_waybel_0(X2,X3)|~l1_orders_2(X3)), inference(csr,[status(thm)],[c_0_33, c_0_18])).
cnf(c_0_42, plain, (r1_lattice3(X1,k6_waybel_0(X2,X3),X4)|r1_orders_2(X2,X3,esk6_3(X1,k6_waybel_0(X2,X3),X4))|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_43, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk6_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_44, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,X1)|v15_waybel_0(esk2_0,esk1_0)|~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_36]), c_0_23])]), c_0_25]), c_0_22])).
cnf(c_0_45, plain, (r2_hidden(X3,k6_waybel_0(X1,X2))|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_46, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk3_0)|v15_waybel_0(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_37, c_0_22])).
cnf(c_0_47, plain, (r1_orders_2(X1,esk7_2(X1,X2),X3)|v1_xboole_0(X2)|v2_struct_0(X1)|~r2_hidden(X3,X2)|~v15_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v13_waybel_0(X2,X1)|~v2_waybel_0(X2,X1)|~l1_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_18]), c_0_40])).
cnf(c_0_48, plain, (r1_lattice3(X1,k6_waybel_0(X2,X3),X4)|r2_hidden(esk6_3(X1,k6_waybel_0(X2,X3),X4),X5)|v2_struct_0(X2)|~r2_hidden(X3,X5)|~m1_subset_1(esk6_3(X1,k6_waybel_0(X2,X3),X4),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,u1_struct_0(X1))|~v13_waybel_0(X5,X2)|~l1_orders_2(X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_18])).
cnf(c_0_49, negated_conjecture, (v13_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_50, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_51, plain, (v15_waybel_0(X3,X2)|v1_xboole_0(X3)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r1_lattice3(X2,X3,X1)|~v2_waybel_0(X3,X2)|~v13_waybel_0(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v3_orders_2(X2)|~v4_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_52, negated_conjecture, (r1_lattice3(esk1_0,X1,esk3_0)|v15_waybel_0(esk2_0,esk1_0)|~r2_hidden(esk6_3(esk1_0,X1,esk3_0),esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_23])]), c_0_22])).
cnf(c_0_53, negated_conjecture, (r2_hidden(esk3_0,k6_waybel_0(esk1_0,esk3_0))|v15_waybel_0(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_23])]), c_0_25]), c_0_22])).
cnf(c_0_54, plain, (r1_lattice3(X1,X2,esk7_2(X1,X3))|v1_xboole_0(X3)|v2_struct_0(X1)|~r2_hidden(esk6_3(X1,X2,esk7_2(X1,X3)),X3)|~v15_waybel_0(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v13_waybel_0(X3,X1)|~v2_waybel_0(X3,X1)|~l1_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_47]), c_0_40])).
cnf(c_0_55, negated_conjecture, (r1_lattice3(X1,k6_waybel_0(esk1_0,X2),X3)|r2_hidden(esk6_3(X1,k6_waybel_0(esk1_0,X2),X3),esk2_0)|~r2_hidden(X2,esk2_0)|~m1_subset_1(esk6_3(X1,k6_waybel_0(esk1_0,X2),X3),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50]), c_0_23])]), c_0_25])).
cnf(c_0_56, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_57, plain, (v15_waybel_0(X1,X2)|v1_xboole_0(X1)|v2_struct_0(X2)|~r1_lattice3(X2,X1,X3)|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v13_waybel_0(X1,X2)|~v2_waybel_0(X1,X2)|~l1_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[c_0_51, c_0_18])).
cnf(c_0_58, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk3_0)|v15_waybel_0(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_35]), c_0_23])]), c_0_22])).
cnf(c_0_59, negated_conjecture, (v2_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_60, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_61, negated_conjecture, (r2_hidden(esk3_0,esk2_0)|v15_waybel_0(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_53, c_0_36])).
cnf(c_0_62, negated_conjecture, (r1_lattice3(X1,k6_waybel_0(esk1_0,X2),esk7_2(X1,esk2_0))|v2_struct_0(X1)|~r2_hidden(X2,esk2_0)|~v15_waybel_0(esk2_0,X1)|~m1_subset_1(esk6_3(X1,k6_waybel_0(esk1_0,X2),esk7_2(X1,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk7_2(X1,esk2_0),u1_struct_0(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X1)))|~v13_waybel_0(esk2_0,X1)|~v2_waybel_0(esk2_0,X1)|~l1_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56])).
cnf(c_0_63, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_64, negated_conjecture, (v15_waybel_0(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_50]), c_0_49]), c_0_59]), c_0_23]), c_0_60]), c_0_24])]), c_0_56]), c_0_25]), c_0_61])).
fof(c_0_65, plain, ![X38, X39]:(v2_struct_0(X38)|(~v3_orders_2(X38)|~l1_orders_2(X38))|(~m1_subset_1(X39,u1_struct_0(X38))|r1_orders_2(X38,X39,X39))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
cnf(c_0_66, negated_conjecture, (r1_lattice3(esk1_0,k6_waybel_0(esk1_0,X1),esk7_2(esk1_0,esk2_0))|~r2_hidden(X1,esk2_0)|~m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64]), c_0_50]), c_0_49]), c_0_59]), c_0_23]), c_0_60]), c_0_24])]), c_0_25])).
cnf(c_0_67, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_68, negated_conjecture, (r1_orders_2(esk1_0,esk7_2(esk1_0,esk2_0),X1)|~r2_hidden(X1,k6_waybel_0(esk1_0,X2))|~r2_hidden(X2,esk2_0)|~m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_66]), c_0_23])])).
cnf(c_0_69, plain, (r2_hidden(X1,k6_waybel_0(X2,X1))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)), inference(spm,[status(thm)],[c_0_45, c_0_67])).
cnf(c_0_70, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_18, c_0_50])).
fof(c_0_71, plain, ![X40, X41, X42]:((m1_subset_1(esk8_3(X40,X41,X42),X40)|r1_tarski(X41,X42)|~m1_subset_1(X42,k1_zfmisc_1(X40))|~m1_subset_1(X41,k1_zfmisc_1(X40)))&((r2_hidden(esk8_3(X40,X41,X42),X41)|r1_tarski(X41,X42)|~m1_subset_1(X42,k1_zfmisc_1(X40))|~m1_subset_1(X41,k1_zfmisc_1(X40)))&(~r2_hidden(esk8_3(X40,X41,X42),X42)|r1_tarski(X41,X42)|~m1_subset_1(X42,k1_zfmisc_1(X40))|~m1_subset_1(X41,k1_zfmisc_1(X40))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_subset_1])])])])])).
cnf(c_0_72, negated_conjecture, (r1_orders_2(esk1_0,esk7_2(esk1_0,esk2_0),X1)|~r2_hidden(X1,esk2_0)|~m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_23]), c_0_24])]), c_0_25]), c_0_70])).
cnf(c_0_73, plain, (r2_hidden(esk8_3(X1,X2,X3),X2)|r1_tarski(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_74, plain, (r1_tarski(X2,X3)|~r2_hidden(esk8_3(X1,X2,X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_75, negated_conjecture, (r2_hidden(X1,k6_waybel_0(esk1_0,esk7_2(esk1_0,esk2_0)))|~r2_hidden(X1,esk2_0)|~m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_72]), c_0_23])]), c_0_25]), c_0_70])).
cnf(c_0_76, plain, (r1_tarski(k6_waybel_0(X1,X2),X3)|r1_orders_2(X1,X2,esk8_3(X4,k6_waybel_0(X1,X2),X3))|v2_struct_0(X1)|~m1_subset_1(k6_waybel_0(X1,X2),k1_zfmisc_1(X4))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(X4))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_34, c_0_73])).
cnf(c_0_77, negated_conjecture, (r1_tarski(X1,k6_waybel_0(esk1_0,esk7_2(esk1_0,esk2_0)))|~r2_hidden(esk8_3(X2,X1,k6_waybel_0(esk1_0,esk7_2(esk1_0,esk2_0))),esk2_0)|~m1_subset_1(k6_waybel_0(esk1_0,esk7_2(esk1_0,esk2_0)),k1_zfmisc_1(X2))|~m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_78, plain, (r1_tarski(k6_waybel_0(X1,X2),X3)|r2_hidden(esk8_3(X4,k6_waybel_0(X1,X2),X3),X5)|v2_struct_0(X1)|~r2_hidden(X2,X5)|~m1_subset_1(esk8_3(X4,k6_waybel_0(X1,X2),X3),u1_struct_0(X1))|~m1_subset_1(k6_waybel_0(X1,X2),k1_zfmisc_1(X4))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(X4))|~v13_waybel_0(X5,X1)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_76]), c_0_18])).
fof(c_0_79, plain, ![X44, X45]:(((r1_tarski(X44,X45)|X44!=X45)&(r1_tarski(X45,X44)|X44!=X45))&(~r1_tarski(X44,X45)|~r1_tarski(X45,X44)|X44=X45)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_80, negated_conjecture, (r1_tarski(esk2_0,k6_waybel_0(esk1_0,esk7_2(esk1_0,esk2_0)))|~m1_subset_1(k6_waybel_0(esk1_0,esk7_2(esk1_0,esk2_0)),k1_zfmisc_1(X1))|~m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_77, c_0_73])).
cnf(c_0_81, negated_conjecture, (~v15_waybel_0(esk2_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|esk2_0!=k6_waybel_0(esk1_0,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_82, negated_conjecture, (r1_tarski(k6_waybel_0(esk1_0,X1),X2)|r2_hidden(esk8_3(X3,k6_waybel_0(esk1_0,X1),X2),esk2_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(esk8_3(X3,k6_waybel_0(esk1_0,X1),X2),u1_struct_0(esk1_0))|~m1_subset_1(k6_waybel_0(esk1_0,X1),k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_49]), c_0_50]), c_0_23])]), c_0_25])).
cnf(c_0_83, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_84, negated_conjecture, (r1_tarski(esk2_0,k6_waybel_0(esk1_0,esk7_2(esk1_0,esk2_0)))|~m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_19]), c_0_50]), c_0_23])]), c_0_25])).
cnf(c_0_85, negated_conjecture, (k6_waybel_0(esk1_0,X1)!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_64])])).
cnf(c_0_86, negated_conjecture, (r1_tarski(k6_waybel_0(esk1_0,X1),esk2_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(esk8_3(X2,k6_waybel_0(esk1_0,X1),esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(k6_waybel_0(esk1_0,X1),k1_zfmisc_1(X2))|~m1_subset_1(esk2_0,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_74, c_0_82])).
cnf(c_0_87, plain, (m1_subset_1(esk8_3(X1,X2,X3),X1)|r1_tarski(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_88, negated_conjecture, (~r1_tarski(k6_waybel_0(esk1_0,esk7_2(esk1_0,esk2_0)),esk2_0)|~m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_85])).
cnf(c_0_89, negated_conjecture, (r1_tarski(k6_waybel_0(esk1_0,X1),esk2_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(k6_waybel_0(esk1_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_50])])).
cnf(c_0_90, negated_conjecture, (~r2_hidden(esk7_2(esk1_0,esk2_0),esk2_0)|~m1_subset_1(k6_waybel_0(esk1_0,esk7_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_70])).
cnf(c_0_91, negated_conjecture, (~r2_hidden(esk7_2(esk1_0,esk2_0),esk2_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_19]), c_0_23])]), c_0_25]), c_0_70])).
cnf(c_0_92, plain, (r2_hidden(esk7_2(X1,X2),X2)|v1_xboole_0(X2)|v2_struct_0(X1)|~v15_waybel_0(X2,X1)|~v2_waybel_0(X2,X1)|~v13_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_93, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_64]), c_0_50]), c_0_49]), c_0_59]), c_0_23]), c_0_60]), c_0_24])]), c_0_56]), c_0_25]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 94
# Proof object clause steps            : 69
# Proof object formula steps           : 25
# Proof object conjectures             : 41
# Proof object clause conjectures      : 38
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 34
# Proof object simplifying inferences  : 93
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 39
# Processed clauses                    : 254
# ...of these trivial                  : 0
# ...subsumed                          : 58
# ...remaining for further processing  : 196
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 14
# Backward-rewritten                   : 33
# Generated clauses                    : 360
# ...of the previous two non-trivial   : 325
# Contextual simplify-reflections      : 44
# Paramodulations                      : 358
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 109
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 98
# Current number of unprocessed clauses: 72
# ...number of literals in the above   : 972
# Current number of archived formulas  : 0
# Current number of archived clauses   : 85
# Clause-clause subsumption calls (NU) : 7479
# Rec. Clause-clause subsumption calls : 630
# Non-unit clause-clause subsumptions  : 116
# Unit Clause-clause subsumption calls : 49
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 18706

# -------------------------------------------------
# User time                : 0.024 s
# System time              : 0.000 s
# Total time               : 0.024 s
# Maximum resident set size: 28832 pages
