# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C41_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.007 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d3_xboole_0, axiom, ![X1, X2, X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2)))), file('data/MPTP/proofs/0ebffd2ff7a7453c81777911e3698c09.E_input', d3_xboole_0)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/0ebffd2ff7a7453c81777911e3698c09.E_input', d4_xboole_0)).
fof(t22_xboole_1, conjecture, ![X1, X2]:k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/0ebffd2ff7a7453c81777911e3698c09.E_input', t22_xboole_1)).
fof(c_0_3, plain, ![X16, X17, X18, X19, X20, X21, X22, X23]:(((~r2_hidden(X19,X18)|(r2_hidden(X19,X16)|r2_hidden(X19,X17))|X18!=k2_xboole_0(X16,X17))&((~r2_hidden(X20,X16)|r2_hidden(X20,X18)|X18!=k2_xboole_0(X16,X17))&(~r2_hidden(X20,X17)|r2_hidden(X20,X18)|X18!=k2_xboole_0(X16,X17))))&(((~r2_hidden(esk4_3(X21,X22,X23),X21)|~r2_hidden(esk4_3(X21,X22,X23),X23)|X23=k2_xboole_0(X21,X22))&(~r2_hidden(esk4_3(X21,X22,X23),X22)|~r2_hidden(esk4_3(X21,X22,X23),X23)|X23=k2_xboole_0(X21,X22)))&(r2_hidden(esk4_3(X21,X22,X23),X23)|(r2_hidden(esk4_3(X21,X22,X23),X21)|r2_hidden(esk4_3(X21,X22,X23),X22))|X23=k2_xboole_0(X21,X22)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])).
fof(c_0_4, plain, ![X7, X8, X9, X10, X11, X12, X13, X14]:((((r2_hidden(X10,X7)|~r2_hidden(X10,X9)|X9!=k3_xboole_0(X7,X8))&(r2_hidden(X10,X8)|~r2_hidden(X10,X9)|X9!=k3_xboole_0(X7,X8)))&(~r2_hidden(X11,X7)|~r2_hidden(X11,X8)|r2_hidden(X11,X9)|X9!=k3_xboole_0(X7,X8)))&((~r2_hidden(esk3_3(X12,X13,X14),X14)|(~r2_hidden(esk3_3(X12,X13,X14),X12)|~r2_hidden(esk3_3(X12,X13,X14),X13))|X14=k3_xboole_0(X12,X13))&((r2_hidden(esk3_3(X12,X13,X14),X12)|r2_hidden(esk3_3(X12,X13,X14),X14)|X14=k3_xboole_0(X12,X13))&(r2_hidden(esk3_3(X12,X13,X14),X13)|r2_hidden(esk3_3(X12,X13,X14),X14)|X14=k3_xboole_0(X12,X13))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_5, plain, (r2_hidden(esk4_3(X1,X2,X3),X3)|r2_hidden(esk4_3(X1,X2,X3),X1)|r2_hidden(esk4_3(X1,X2,X3),X2)|X3=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_6, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, plain, (X3=k2_xboole_0(X1,X2)|~r2_hidden(esk4_3(X1,X2,X3),X1)|~r2_hidden(esk4_3(X1,X2,X3),X3)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_8, plain, (k2_xboole_0(X1,X2)=X1|r2_hidden(esk4_3(X1,X2,X1),X1)|r2_hidden(esk4_3(X1,X2,X1),X2)), inference(ef,[status(thm)],[c_0_5])).
fof(c_0_9, negated_conjecture, ~(![X1, X2]:k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(assume_negation,[status(cth)],[t22_xboole_1])).
cnf(c_0_10, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_6])).
cnf(c_0_11, plain, (k2_xboole_0(X1,X2)=X1|r2_hidden(esk4_3(X1,X2,X1),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_8])).
fof(c_0_12, negated_conjecture, k2_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0))!=esk1_0, inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
cnf(c_0_13, plain, (k2_xboole_0(X1,k3_xboole_0(X2,X3))=X1|r2_hidden(esk4_3(X1,k3_xboole_0(X2,X3),X1),X2)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_14, negated_conjecture, (k2_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0))!=esk1_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_13]), c_0_13])).
cnf(c_0_16, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_15])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 17
# Proof object clause steps            : 10
# Proof object formula steps           : 7
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 4
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 4
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 13
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 13
# Processed clauses                    : 140
# ...of these trivial                  : 19
# ...subsumed                          : 62
# ...remaining for further processing  : 59
# Other redundant clauses eliminated   : 24
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 3
# Generated clauses                    : 614
# ...of the previous two non-trivial   : 422
# Contextual simplify-reflections      : 11
# Paramodulations                      : 546
# Factorizations                       : 38
# Equation resolutions                 : 30
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 54
#    Positive orientable unit clauses  : 15
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 39
# Current number of unprocessed clauses: 291
# ...number of literals in the above   : 944
# Current number of archived formulas  : 0
# Current number of archived clauses   : 5
# Clause-clause subsumption calls (NU) : 500
# Rec. Clause-clause subsumption calls : 326
# Non-unit clause-clause subsumptions  : 75
# Unit Clause-clause subsumption calls : 101
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 22
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6133

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.003 s
# Total time               : 0.011 s
# Maximum resident set size: 28688 pages
