# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/0ebf44e8209d474e9d025634da22fbd7.E_input', d10_xboole_0)).
fof(d1_wellord2, axiom, ![X1, X2]:(v1_relat_1(X2)=>(X2=k1_wellord2(X1)<=>(k3_relat_1(X2)=X1&![X3, X4]:((r2_hidden(X3,X1)&r2_hidden(X4,X1))=>(r2_hidden(k4_tarski(X3,X4),X2)<=>r1_tarski(X3,X4)))))), file('data/MPTP/proofs/0ebf44e8209d474e9d025634da22fbd7.E_input', d1_wellord2)).
fof(dt_k1_wellord2, axiom, ![X1]:v1_relat_1(k1_wellord2(X1)), file('data/MPTP/proofs/0ebf44e8209d474e9d025634da22fbd7.E_input', dt_k1_wellord2)).
fof(t36_wellord1, axiom, ![X1, X2]:(v1_relat_1(X2)=>((v2_wellord1(X2)&r1_tarski(X1,k3_relat_1(X2)))=>(~((~(X1=k3_relat_1(X2))&![X3]:~((r2_hidden(X3,k3_relat_1(X2))&X1=k1_wellord1(X2,X3)))))<=>![X3]:(r2_hidden(X3,X1)=>![X4]:(r2_hidden(k4_tarski(X4,X3),X2)=>r2_hidden(X4,X1)))))), file('data/MPTP/proofs/0ebf44e8209d474e9d025634da22fbd7.E_input', t36_wellord1)).
fof(d1_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(X3=k1_wellord1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(~(X4=X2)&r2_hidden(k4_tarski(X4,X2),X1))))), file('data/MPTP/proofs/0ebf44e8209d474e9d025634da22fbd7.E_input', d1_wellord1)).
fof(t57_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v2_wellord1(X1)=>![X2]:~((r2_hidden(X2,k3_relat_1(X1))&r4_wellord1(X1,k2_wellord1(X1,k1_wellord1(X1,X2))))))), file('data/MPTP/proofs/0ebf44e8209d474e9d025634da22fbd7.E_input', t57_wellord1)).
fof(t10_wellord2, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)=>X1=k1_wellord1(k1_wellord2(X2),X1)))), file('data/MPTP/proofs/0ebf44e8209d474e9d025634da22fbd7.E_input', t10_wellord2)).
fof(t7_wellord2, axiom, ![X1]:(v3_ordinal1(X1)=>v2_wellord1(k1_wellord2(X1))), file('data/MPTP/proofs/0ebf44e8209d474e9d025634da22fbd7.E_input', t7_wellord2)).
fof(t11_wellord2, conjecture, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r4_wellord1(k1_wellord2(X1),k1_wellord2(X2))=>X1=X2))), file('data/MPTP/proofs/0ebf44e8209d474e9d025634da22fbd7.E_input', t11_wellord2)).
fof(t8_wellord2, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_wellord1(k1_wellord2(X2),X1)=k1_wellord2(X1)), file('data/MPTP/proofs/0ebf44e8209d474e9d025634da22fbd7.E_input', t8_wellord2)).
fof(t50_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r4_wellord1(X1,X2)=>r4_wellord1(X2,X1)))), file('data/MPTP/proofs/0ebf44e8209d474e9d025634da22fbd7.E_input', t50_wellord1)).
fof(t24_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>~((~(r2_hidden(X1,X2))&(~(X1=X2)&~(r2_hidden(X2,X1))))))), file('data/MPTP/proofs/0ebf44e8209d474e9d025634da22fbd7.E_input', t24_ordinal1)).
fof(t41_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v2_wellord1(X1)=>![X2]:(![X3]:((r2_hidden(X3,k3_relat_1(X1))&r1_tarski(k1_wellord1(X1,X3),X2))=>r2_hidden(X3,X2))=>r1_tarski(k3_relat_1(X1),X2)))), file('data/MPTP/proofs/0ebf44e8209d474e9d025634da22fbd7.E_input', t41_wellord1)).
fof(c_0_13, plain, ![X9, X10]:(((r1_tarski(X9,X10)|X9!=X10)&(r1_tarski(X10,X9)|X9!=X10))&(~r1_tarski(X9,X10)|~r1_tarski(X10,X9)|X9=X10)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_14, plain, ![X18, X19, X20, X21]:(((k3_relat_1(X19)=X18|X19!=k1_wellord2(X18)|~v1_relat_1(X19))&((~r2_hidden(k4_tarski(X20,X21),X19)|r1_tarski(X20,X21)|(~r2_hidden(X20,X18)|~r2_hidden(X21,X18))|X19!=k1_wellord2(X18)|~v1_relat_1(X19))&(~r1_tarski(X20,X21)|r2_hidden(k4_tarski(X20,X21),X19)|(~r2_hidden(X20,X18)|~r2_hidden(X21,X18))|X19!=k1_wellord2(X18)|~v1_relat_1(X19))))&(((r2_hidden(esk4_2(X18,X19),X18)|k3_relat_1(X19)!=X18|X19=k1_wellord2(X18)|~v1_relat_1(X19))&(r2_hidden(esk5_2(X18,X19),X18)|k3_relat_1(X19)!=X18|X19=k1_wellord2(X18)|~v1_relat_1(X19)))&((~r2_hidden(k4_tarski(esk4_2(X18,X19),esk5_2(X18,X19)),X19)|~r1_tarski(esk4_2(X18,X19),esk5_2(X18,X19))|k3_relat_1(X19)!=X18|X19=k1_wellord2(X18)|~v1_relat_1(X19))&(r2_hidden(k4_tarski(esk4_2(X18,X19),esk5_2(X18,X19)),X19)|r1_tarski(esk4_2(X18,X19),esk5_2(X18,X19))|k3_relat_1(X19)!=X18|X19=k1_wellord2(X18)|~v1_relat_1(X19))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_wellord2])])])])])).
fof(c_0_15, plain, ![X45]:v1_relat_1(k1_wellord2(X45)), inference(variable_rename,[status(thm)],[dt_k1_wellord2])).
fof(c_0_16, plain, ![X37, X38, X39, X40, X41]:(((X37!=k3_relat_1(X38)|(~r2_hidden(X40,X37)|(~r2_hidden(k4_tarski(X41,X40),X38)|r2_hidden(X41,X37)))|(~v2_wellord1(X38)|~r1_tarski(X37,k3_relat_1(X38)))|~v1_relat_1(X38))&(~r2_hidden(X39,k3_relat_1(X38))|X37!=k1_wellord1(X38,X39)|(~r2_hidden(X40,X37)|(~r2_hidden(k4_tarski(X41,X40),X38)|r2_hidden(X41,X37)))|(~v2_wellord1(X38)|~r1_tarski(X37,k3_relat_1(X38)))|~v1_relat_1(X38)))&(((r2_hidden(esk9_2(X37,X38),k3_relat_1(X38))|X37=k3_relat_1(X38)|r2_hidden(esk7_2(X37,X38),X37)|(~v2_wellord1(X38)|~r1_tarski(X37,k3_relat_1(X38)))|~v1_relat_1(X38))&(X37=k1_wellord1(X38,esk9_2(X37,X38))|X37=k3_relat_1(X38)|r2_hidden(esk7_2(X37,X38),X37)|(~v2_wellord1(X38)|~r1_tarski(X37,k3_relat_1(X38)))|~v1_relat_1(X38)))&(((r2_hidden(esk9_2(X37,X38),k3_relat_1(X38))|X37=k3_relat_1(X38)|r2_hidden(k4_tarski(esk8_2(X37,X38),esk7_2(X37,X38)),X38)|(~v2_wellord1(X38)|~r1_tarski(X37,k3_relat_1(X38)))|~v1_relat_1(X38))&(X37=k1_wellord1(X38,esk9_2(X37,X38))|X37=k3_relat_1(X38)|r2_hidden(k4_tarski(esk8_2(X37,X38),esk7_2(X37,X38)),X38)|(~v2_wellord1(X38)|~r1_tarski(X37,k3_relat_1(X38)))|~v1_relat_1(X38)))&((r2_hidden(esk9_2(X37,X38),k3_relat_1(X38))|X37=k3_relat_1(X38)|~r2_hidden(esk8_2(X37,X38),X37)|(~v2_wellord1(X38)|~r1_tarski(X37,k3_relat_1(X38)))|~v1_relat_1(X38))&(X37=k1_wellord1(X38,esk9_2(X37,X38))|X37=k3_relat_1(X38)|~r2_hidden(esk8_2(X37,X38),X37)|(~v2_wellord1(X38)|~r1_tarski(X37,k3_relat_1(X38)))|~v1_relat_1(X38)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t36_wellord1])])])])])])).
cnf(c_0_17, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_18, plain, ![X24, X25, X26, X27, X28, X29, X30]:((((X27!=X25|~r2_hidden(X27,X26)|X26!=k1_wellord1(X24,X25)|~v1_relat_1(X24))&(r2_hidden(k4_tarski(X27,X25),X24)|~r2_hidden(X27,X26)|X26!=k1_wellord1(X24,X25)|~v1_relat_1(X24)))&(X28=X25|~r2_hidden(k4_tarski(X28,X25),X24)|r2_hidden(X28,X26)|X26!=k1_wellord1(X24,X25)|~v1_relat_1(X24)))&((~r2_hidden(esk6_3(X24,X29,X30),X30)|(esk6_3(X24,X29,X30)=X29|~r2_hidden(k4_tarski(esk6_3(X24,X29,X30),X29),X24))|X30=k1_wellord1(X24,X29)|~v1_relat_1(X24))&((esk6_3(X24,X29,X30)!=X29|r2_hidden(esk6_3(X24,X29,X30),X30)|X30=k1_wellord1(X24,X29)|~v1_relat_1(X24))&(r2_hidden(k4_tarski(esk6_3(X24,X29,X30),X29),X24)|r2_hidden(esk6_3(X24,X29,X30),X30)|X30=k1_wellord1(X24,X29)|~v1_relat_1(X24))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_wellord1])])])])])])])).
fof(c_0_19, plain, ![X7, X8]:(~v1_relat_1(X7)|(~v2_wellord1(X7)|(~r2_hidden(X8,k3_relat_1(X7))|~r4_wellord1(X7,k2_wellord1(X7,k1_wellord1(X7,X8)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t57_wellord1])])])).
fof(c_0_20, plain, ![X13, X14]:(~v3_ordinal1(X13)|(~v3_ordinal1(X14)|(~r2_hidden(X13,X14)|X13=k1_wellord1(k1_wellord2(X14),X13)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_wellord2])])])).
cnf(c_0_21, plain, (k3_relat_1(X1)=X2|X1!=k1_wellord2(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (v1_relat_1(k1_wellord2(X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_23, plain, ![X32]:(~v3_ordinal1(X32)|v2_wellord1(k1_wellord2(X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_wellord2])])).
fof(c_0_24, negated_conjecture, ~(![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r4_wellord1(k1_wellord2(X1),k1_wellord2(X2))=>X1=X2)))), inference(assume_negation,[status(cth)],[t11_wellord2])).
cnf(c_0_25, plain, (r2_hidden(X4,X1)|X1!=k3_relat_1(X2)|~r2_hidden(X3,X1)|~r2_hidden(k4_tarski(X4,X3),X2)|~v2_wellord1(X2)|~r1_tarski(X1,k3_relat_1(X2))|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_27, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,X4)|X4!=k1_wellord1(X3,X2)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, plain, (~v1_relat_1(X1)|~v2_wellord1(X1)|~r2_hidden(X2,k3_relat_1(X1))|~r4_wellord1(X1,k2_wellord1(X1,k1_wellord1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (X1=k1_wellord1(k1_wellord2(X2),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30, plain, (k3_relat_1(k1_wellord2(X1))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_21]), c_0_22])])).
cnf(c_0_31, plain, (v2_wellord1(k1_wellord2(X1))|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_32, plain, ![X33, X34]:(~r1_tarski(X33,X34)|k2_wellord1(k1_wellord2(X34),X33)=k1_wellord2(X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_wellord2])])).
fof(c_0_33, plain, ![X35, X36]:(~v1_relat_1(X35)|(~v1_relat_1(X36)|(~r4_wellord1(X35,X36)|r4_wellord1(X36,X35)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t50_wellord1])])])).
fof(c_0_34, negated_conjecture, (v3_ordinal1(esk1_0)&(v3_ordinal1(esk2_0)&(r4_wellord1(k1_wellord2(esk1_0),k1_wellord2(esk2_0))&esk1_0!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).
fof(c_0_35, plain, ![X11, X12]:(~v3_ordinal1(X11)|(~v3_ordinal1(X12)|(r2_hidden(X11,X12)|(X11=X12|r2_hidden(X12,X11))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_ordinal1])])])])).
cnf(c_0_36, plain, (r2_hidden(X1,k3_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)|~r2_hidden(X3,k3_relat_1(X2))|~v2_wellord1(X2)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_25]), c_0_26])])).
cnf(c_0_37, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,k1_wellord1(X3,X2))|~v1_relat_1(X3)), inference(er,[status(thm)],[c_0_27])).
fof(c_0_38, plain, ![X15, X16]:(((r2_hidden(esk3_2(X15,X16),k3_relat_1(X15))|r1_tarski(k3_relat_1(X15),X16)|~v2_wellord1(X15)|~v1_relat_1(X15))&(r1_tarski(k1_wellord1(X15,esk3_2(X15,X16)),X16)|r1_tarski(k3_relat_1(X15),X16)|~v2_wellord1(X15)|~v1_relat_1(X15)))&(~r2_hidden(esk3_2(X15,X16),X16)|r1_tarski(k3_relat_1(X15),X16)|~v2_wellord1(X15)|~v1_relat_1(X15))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t41_wellord1])])])])])).
cnf(c_0_39, plain, (~r2_hidden(X1,X2)|~r4_wellord1(k1_wellord2(X2),k2_wellord1(k1_wellord2(X2),X1))|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_22])]), c_0_30])]), c_0_31])).
cnf(c_0_40, plain, (k2_wellord1(k1_wellord2(X2),X1)=k1_wellord2(X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_41, plain, (r4_wellord1(X2,X1)|~v1_relat_1(X1)|~v1_relat_1(X2)|~r4_wellord1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, negated_conjecture, (r4_wellord1(k1_wellord2(esk1_0),k1_wellord2(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43, plain, (r2_hidden(X1,X2)|X1=X2|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_44, negated_conjecture, (v3_ordinal1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45, plain, (r2_hidden(X1,k3_relat_1(X2))|~r2_hidden(X1,k1_wellord1(X2,X3))|~r2_hidden(X3,k3_relat_1(X2))|~v2_wellord1(X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_46, plain, (r2_hidden(esk3_2(X1,X2),k3_relat_1(X1))|r1_tarski(k3_relat_1(X1),X2)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_47, plain, (~r1_tarski(X1,X2)|~r2_hidden(X1,X2)|~r4_wellord1(k1_wellord2(X2),k1_wellord2(X1))|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_48, negated_conjecture, (r4_wellord1(k1_wellord2(esk2_0),k1_wellord2(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_22]), c_0_22])])).
cnf(c_0_49, negated_conjecture, (v3_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_50, negated_conjecture, (X1=esk2_0|r2_hidden(esk2_0,X1)|r2_hidden(X1,esk2_0)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_51, negated_conjecture, (esk1_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_52, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r2_hidden(X3,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X3)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_29]), c_0_30]), c_0_30]), c_0_22])]), c_0_31])).
cnf(c_0_53, plain, (r1_tarski(X1,X2)|r2_hidden(esk3_2(k1_wellord2(X1),X2),X1)|~v2_wellord1(k1_wellord2(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_30]), c_0_22])])).
cnf(c_0_54, negated_conjecture, (~r1_tarski(esk1_0,esk2_0)|~r2_hidden(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_44]), c_0_49])])).
cnf(c_0_55, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|r2_hidden(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_49]), c_0_51])).
cnf(c_0_56, plain, (r1_tarski(k3_relat_1(X1),X2)|~r2_hidden(esk3_2(X1,X2),X2)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_57, plain, (r1_tarski(X1,X2)|r2_hidden(esk3_2(k1_wellord2(X1),X2),X3)|~r2_hidden(X1,X3)|~v3_ordinal1(X3)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_31])).
cnf(c_0_58, negated_conjecture, (~r1_tarski(esk2_0,esk1_0)|~r2_hidden(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_42]), c_0_49]), c_0_44])])).
cnf(c_0_59, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r1_tarski(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_60, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_30]), c_0_22])]), c_0_31])).
cnf(c_0_61, negated_conjecture, (~r1_tarski(esk2_0,esk1_0)|~r1_tarski(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_62, negated_conjecture, (~r1_tarski(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_59]), c_0_49]), c_0_44])]), c_0_61])).
cnf(c_0_63, negated_conjecture, (r2_hidden(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_55]), c_0_44]), c_0_49])]), c_0_62])).
cnf(c_0_64, negated_conjecture, (~r1_tarski(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_63])])).
cnf(c_0_65, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_63]), c_0_49]), c_0_44])]), c_0_64]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 66
# Proof object clause steps            : 39
# Proof object formula steps           : 27
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 49
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 38
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 38
# Processed clauses                    : 99
# ...of these trivial                  : 2
# ...subsumed                          : 9
# ...remaining for further processing  : 87
# Other redundant clauses eliminated   : 15
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 3
# Generated clauses                    : 150
# ...of the previous two non-trivial   : 142
# Contextual simplify-reflections      : 12
# Paramodulations                      : 136
# Factorizations                       : 0
# Equation resolutions                 : 15
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 65
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 53
# Current number of unprocessed clauses: 81
# ...number of literals in the above   : 546
# Current number of archived formulas  : 0
# Current number of archived clauses   : 8
# Clause-clause subsumption calls (NU) : 716
# Rec. Clause-clause subsumption calls : 126
# Non-unit clause-clause subsumptions  : 24
# Unit Clause-clause subsumption calls : 114
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6298

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.000 s
# Total time               : 0.013 s
# Maximum resident set size: 28748 pages
