# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(fc1_relat_1, axiom, ![X1, X2]:(v1_relat_1(X1)=>v1_relat_1(k3_xboole_0(X1,X2))), file('data/MPTP/proofs/0f805d5bd4994f209ac8ec998bf5f2aa.E_input', fc1_relat_1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/0f805d5bd4994f209ac8ec998bf5f2aa.E_input', t12_setfam_1)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/0f805d5bd4994f209ac8ec998bf5f2aa.E_input', t2_boole)).
fof(t152_zfmisc_1, axiom, ![X1, X2]:~(r2_hidden(X1,k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/0f805d5bd4994f209ac8ec998bf5f2aa.E_input', t152_zfmisc_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/0f805d5bd4994f209ac8ec998bf5f2aa.E_input', t113_zfmisc_1)).
fof(t62_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>(k5_relat_1(k1_xboole_0,X1)=k1_xboole_0&k5_relat_1(X1,k1_xboole_0)=k1_xboole_0)), file('data/MPTP/proofs/0f805d5bd4994f209ac8ec998bf5f2aa.E_input', t62_relat_1)).
fof(d8_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(X3=k5_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>?[X6]:(r2_hidden(k4_tarski(X4,X6),X1)&r2_hidden(k4_tarski(X6,X5),X2))))))), file('data/MPTP/proofs/0f805d5bd4994f209ac8ec998bf5f2aa.E_input', d8_relat_1)).
fof(c_0_7, plain, ![X8, X9]:(~v1_relat_1(X8)|v1_relat_1(k3_xboole_0(X8,X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_relat_1])])).
fof(c_0_8, plain, ![X10, X11]:k1_setfam_1(k2_tarski(X10,X11))=k3_xboole_0(X10,X11), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_9, plain, ![X25]:k3_xboole_0(X25,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_10, plain, ![X28, X29]:~r2_hidden(X28,k2_zfmisc_1(X28,X29)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_zfmisc_1])])).
fof(c_0_11, plain, ![X26, X27]:((k2_zfmisc_1(X26,X27)!=k1_xboole_0|(X26=k1_xboole_0|X27=k1_xboole_0))&((X26!=k1_xboole_0|k2_zfmisc_1(X26,X27)=k1_xboole_0)&(X27!=k1_xboole_0|k2_zfmisc_1(X26,X27)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_12, plain, (~epred1_0<=>![X2]:X2!=k1_xboole_0), introduced(definition)).
cnf(c_0_13, plain, (v1_relat_1(k3_xboole_0(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_16, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>(k5_relat_1(k1_xboole_0,X1)=k1_xboole_0&k5_relat_1(X1,k1_xboole_0)=k1_xboole_0))), inference(assume_negation,[status(cth)],[t62_relat_1])).
fof(c_0_17, plain, (~epred2_0<=>![X1]:~r2_hidden(X1,k1_xboole_0)), introduced(definition)).
cnf(c_0_18, plain, (~r2_hidden(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20, plain, (epred1_0|X1!=k1_xboole_0), inference(split_equiv,[status(thm)],[c_0_12])).
cnf(c_0_21, plain, (v1_relat_1(k1_setfam_1(k2_tarski(X1,X2)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_22, plain, (k1_setfam_1(k2_tarski(X1,k1_xboole_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_15, c_0_14])).
fof(c_0_23, negated_conjecture, (v1_relat_1(esk1_0)&(k5_relat_1(k1_xboole_0,esk1_0)!=k1_xboole_0|k5_relat_1(esk1_0,k1_xboole_0)!=k1_xboole_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).
cnf(c_0_24, plain, (~epred2_0|~epred1_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_12]), c_0_17])).
cnf(c_0_25, plain, (epred1_0), inference(er,[status(thm)],[c_0_20])).
fof(c_0_26, plain, ![X12, X13, X14, X15, X16, X18, X19, X20, X23]:((((r2_hidden(k4_tarski(X15,esk2_5(X12,X13,X14,X15,X16)),X12)|~r2_hidden(k4_tarski(X15,X16),X14)|X14!=k5_relat_1(X12,X13)|~v1_relat_1(X14)|~v1_relat_1(X13)|~v1_relat_1(X12))&(r2_hidden(k4_tarski(esk2_5(X12,X13,X14,X15,X16),X16),X13)|~r2_hidden(k4_tarski(X15,X16),X14)|X14!=k5_relat_1(X12,X13)|~v1_relat_1(X14)|~v1_relat_1(X13)|~v1_relat_1(X12)))&(~r2_hidden(k4_tarski(X18,X20),X12)|~r2_hidden(k4_tarski(X20,X19),X13)|r2_hidden(k4_tarski(X18,X19),X14)|X14!=k5_relat_1(X12,X13)|~v1_relat_1(X14)|~v1_relat_1(X13)|~v1_relat_1(X12)))&((~r2_hidden(k4_tarski(esk3_3(X12,X13,X14),esk4_3(X12,X13,X14)),X14)|(~r2_hidden(k4_tarski(esk3_3(X12,X13,X14),X23),X12)|~r2_hidden(k4_tarski(X23,esk4_3(X12,X13,X14)),X13))|X14=k5_relat_1(X12,X13)|~v1_relat_1(X14)|~v1_relat_1(X13)|~v1_relat_1(X12))&((r2_hidden(k4_tarski(esk3_3(X12,X13,X14),esk5_3(X12,X13,X14)),X12)|r2_hidden(k4_tarski(esk3_3(X12,X13,X14),esk4_3(X12,X13,X14)),X14)|X14=k5_relat_1(X12,X13)|~v1_relat_1(X14)|~v1_relat_1(X13)|~v1_relat_1(X12))&(r2_hidden(k4_tarski(esk5_3(X12,X13,X14),esk4_3(X12,X13,X14)),X13)|r2_hidden(k4_tarski(esk3_3(X12,X13,X14),esk4_3(X12,X13,X14)),X14)|X14=k5_relat_1(X12,X13)|~v1_relat_1(X14)|~v1_relat_1(X13)|~v1_relat_1(X12))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_relat_1])])])])])])).
cnf(c_0_27, plain, (v1_relat_1(k1_xboole_0)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (~epred2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25])])).
cnf(c_0_30, plain, (r2_hidden(k4_tarski(esk3_3(X1,X2,X3),esk5_3(X1,X2,X3)),X1)|r2_hidden(k4_tarski(esk3_3(X1,X2,X3),esk4_3(X1,X2,X3)),X3)|X3=k5_relat_1(X1,X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_31, negated_conjecture, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_32, plain, (~r2_hidden(X1,k1_xboole_0)), inference(sr,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_17]), c_0_29])).
cnf(c_0_33, negated_conjecture, (k5_relat_1(X1,X2)=k1_xboole_0|r2_hidden(k4_tarski(esk3_3(X1,X2,k1_xboole_0),esk5_3(X1,X2,k1_xboole_0)),X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])).
cnf(c_0_34, plain, (r2_hidden(k4_tarski(esk5_3(X1,X2,X3),esk4_3(X1,X2,X3)),X2)|r2_hidden(k4_tarski(esk3_3(X1,X2,X3),esk4_3(X1,X2,X3)),X3)|X3=k5_relat_1(X1,X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35, negated_conjecture, (k5_relat_1(X1,esk1_0)=k1_xboole_0|r2_hidden(k4_tarski(esk3_3(X1,esk1_0,k1_xboole_0),esk5_3(X1,esk1_0,k1_xboole_0)),X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_33, c_0_28])).
cnf(c_0_36, negated_conjecture, (k5_relat_1(X1,X2)=k1_xboole_0|r2_hidden(k4_tarski(esk5_3(X1,X2,k1_xboole_0),esk4_3(X1,X2,k1_xboole_0)),X2)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_31]), c_0_32])).
cnf(c_0_37, negated_conjecture, (k5_relat_1(k1_xboole_0,esk1_0)!=k1_xboole_0|k5_relat_1(esk1_0,k1_xboole_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_38, negated_conjecture, (k5_relat_1(k1_xboole_0,esk1_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_31]), c_0_32])).
cnf(c_0_39, negated_conjecture, (k5_relat_1(X1,k1_xboole_0)=k1_xboole_0|~v1_relat_1(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_31]), c_0_32])).
cnf(c_0_40, negated_conjecture, (k5_relat_1(esk1_0,k1_xboole_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_38])])).
cnf(c_0_41, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_28]), c_0_40]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 42
# Proof object clause steps            : 25
# Proof object formula steps           : 17
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 14
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 15
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 14
# Processed clauses                    : 82
# ...of these trivial                  : 0
# ...subsumed                          : 13
# ...remaining for further processing  : 69
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 4
# Generated clauses                    : 113
# ...of the previous two non-trivial   : 110
# Contextual simplify-reflections      : 1
# Paramodulations                      : 107
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 47
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 37
# Current number of unprocessed clauses: 55
# ...number of literals in the above   : 197
# Current number of archived formulas  : 0
# Current number of archived clauses   : 22
# Clause-clause subsumption calls (NU) : 767
# Rec. Clause-clause subsumption calls : 253
# Non-unit clause-clause subsumptions  : 15
# Unit Clause-clause subsumption calls : 48
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4193

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28720 pages
