# No SInE strategy applied
# Auto-Mode selected heuristic G_E___200_B02_F1_SE_CS_SP_PI_S2k
# and selection function SelectCQArNTEqFirst.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t76_funct_2, conjecture, ![X1, X2]:((v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>((r2_relset_1(X1,X1,k1_partfun1(X1,X1,X1,X1,X3,X2),X2)&v2_funct_1(X2))=>r2_relset_1(X1,X1,X3,k6_partfun1(X1))))), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', t76_funct_2)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', d10_xboole_0)).
fof(t67_funct_2, axiom, ![X1, X2]:((v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>k1_relset_1(X1,X1,X2)=X1), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', t67_funct_2)).
fof(t4_funct_2, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r1_tarski(k2_relat_1(X2),X1)=>(v1_funct_1(X2)&(v1_funct_2(X2,k1_relat_1(X2),X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X2),X1))))))), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', t4_funct_2)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', redefinition_k1_relset_1)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', fc6_relat_1)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', redefinition_k2_relset_1)).
fof(redefinition_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:(((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&(v1_funct_1(X6)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))))=>k1_partfun1(X1,X2,X3,X4,X5,X6)=k5_relat_1(X5,X6)), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', redefinition_k1_partfun1)).
fof(t31_funct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>((v2_funct_1(X3)&k2_relset_1(X1,X2,X3)=X2)=>(v1_funct_1(k2_funct_1(X3))&(v1_funct_2(k2_funct_1(X3),X2,X1)&m1_subset_1(k2_funct_1(X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1))))))), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', t31_funct_2)).
fof(redefinition_r2_relset_1, axiom, ![X1, X2, X3, X4]:((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>(r2_relset_1(X1,X2,X3,X4)<=>X3=X4)), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', redefinition_r2_relset_1)).
fof(dt_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:(((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&(v1_funct_1(X6)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))))=>(v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))&m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4))))), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', dt_k1_partfun1)).
fof(t61_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)=>(k5_relat_1(X1,k2_funct_1(X1))=k6_relat_1(k1_relat_1(X1))&k5_relat_1(k2_funct_1(X1),X1)=k6_relat_1(k2_relat_1(X1))))), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', t61_funct_1)).
fof(redefinition_k6_partfun1, axiom, ![X1]:k6_partfun1(X1)=k6_relat_1(X1), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', redefinition_k6_partfun1)).
fof(redefinition_k4_relset_1, axiom, ![X1, X2, X3, X4, X5, X6]:((m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>k4_relset_1(X1,X2,X3,X4,X5,X6)=k5_relat_1(X5,X6)), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', redefinition_k4_relset_1)).
fof(dt_k6_partfun1, axiom, ![X1]:(v1_partfun1(k6_partfun1(X1),X1)&m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', dt_k6_partfun1)).
fof(t23_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(r2_relset_1(X1,X2,k4_relset_1(X1,X1,X1,X2,k6_partfun1(X1),X3),X3)&r2_relset_1(X1,X2,k4_relset_1(X1,X2,X2,X2,X3,k6_partfun1(X2)),X3))), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', t23_funct_2)).
fof(cc1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_partfun1(X3,X1))=>(v1_funct_1(X3)&v1_funct_2(X3,X1,X2)))), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', cc1_funct_2)).
fof(t27_funct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&~((v2_funct_1(X3)<=>![X4, X5]:((v1_funct_1(X5)&(v1_funct_2(X5,X4,X1)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X1)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,X4,X1)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X1)))))=>(r2_relset_1(X4,X2,k1_partfun1(X4,X1,X1,X2,X5,X3),k1_partfun1(X4,X1,X1,X2,X6,X3))=>r2_relset_1(X4,X1,X5,X6)))))))))), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', t27_funct_2)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', t3_subset)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', t113_zfmisc_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/0fb6117e4b6747559e7e4b87e8f48b85.E_input', t2_xboole_1)).
fof(c_0_22, negated_conjecture, ~(![X1, X2]:((v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>((r2_relset_1(X1,X1,k1_partfun1(X1,X1,X1,X1,X3,X2),X2)&v2_funct_1(X2))=>r2_relset_1(X1,X1,X3,k6_partfun1(X1)))))), inference(assume_negation,[status(cth)],[t76_funct_2])).
fof(c_0_23, plain, ![X10, X11]:(((r1_tarski(X10,X11)|X10!=X11)&(r1_tarski(X11,X10)|X10!=X11))&(~r1_tarski(X10,X11)|~r1_tarski(X11,X10)|X10=X11)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_24, plain, ![X12, X13]:(~v1_funct_1(X13)|(~v1_funct_2(X13,X12,X12)|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X12,X12))))|k1_relset_1(X12,X12,X13)=X12), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t67_funct_2])])).
fof(c_0_25, negated_conjecture, ((v1_funct_1(esk2_0)&(v1_funct_2(esk2_0,esk1_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,esk1_0,esk1_0)&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))))&((r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0),esk2_0)&v2_funct_1(esk2_0))&~r2_relset_1(esk1_0,esk1_0,esk3_0,k6_partfun1(esk1_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])).
fof(c_0_26, plain, ![X14, X15]:((v1_funct_1(X15)|~r1_tarski(k2_relat_1(X15),X14)|(~v1_relat_1(X15)|~v1_funct_1(X15)))&((v1_funct_2(X15,k1_relat_1(X15),X14)|~r1_tarski(k2_relat_1(X15),X14)|(~v1_relat_1(X15)|~v1_funct_1(X15)))&(m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X15),X14)))|~r1_tarski(k2_relat_1(X15),X14)|(~v1_relat_1(X15)|~v1_funct_1(X15))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_funct_2])])])).
cnf(c_0_27, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_28, plain, ![X16, X17, X18]:(~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(X16,X17)))|k1_relset_1(X16,X17,X18)=k1_relat_1(X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_29, plain, (k1_relset_1(X2,X2,X1)=X2|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, negated_conjecture, (v1_funct_2(esk2_0,esk1_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_33, plain, ![X19, X20]:(~v1_relat_1(X19)|(~m1_subset_1(X20,k1_zfmisc_1(X19))|v1_relat_1(X20))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_34, plain, ![X21, X22]:v1_relat_1(k2_zfmisc_1(X21,X22)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_35, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),X2)))|~r1_tarski(k2_relat_1(X1),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, negated_conjecture, (k1_relset_1(esk1_0,esk1_0,esk2_0)=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_32])])).
cnf(c_0_39, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_40, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_41, plain, ![X23, X24, X25]:(~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(X23,X24)))|k2_relset_1(X23,X24,X25)=k2_relat_1(X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
cnf(c_0_42, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_43, negated_conjecture, (k1_relat_1(esk2_0)=esk1_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_31]), c_0_38])).
cnf(c_0_44, negated_conjecture, (v1_relat_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_31]), c_0_40])])).
cnf(c_0_45, plain, (v1_funct_2(X1,k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_46, plain, ![X26, X27, X28, X29, X30, X31]:(~v1_funct_1(X30)|~m1_subset_1(X30,k1_zfmisc_1(k2_zfmisc_1(X26,X27)))|(~v1_funct_1(X31)|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(X28,X29))))|k1_partfun1(X26,X27,X28,X29,X30,X31)=k5_relat_1(X30,X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_partfun1])])).
fof(c_0_47, plain, ![X32, X33, X34]:((v1_funct_1(k2_funct_1(X34))|(~v2_funct_1(X34)|k2_relset_1(X32,X33,X34)!=X33)|(~v1_funct_1(X34)|(~v1_funct_2(X34,X32,X33)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))))&((v1_funct_2(k2_funct_1(X34),X33,X32)|(~v2_funct_1(X34)|k2_relset_1(X32,X33,X34)!=X33)|(~v1_funct_1(X34)|(~v1_funct_2(X34,X32,X33)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))))&(m1_subset_1(k2_funct_1(X34),k1_zfmisc_1(k2_zfmisc_1(X33,X32)))|(~v2_funct_1(X34)|k2_relset_1(X32,X33,X34)!=X33)|(~v1_funct_1(X34)|(~v1_funct_2(X34,X32,X33)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_funct_2])])])).
cnf(c_0_48, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_49, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k2_relat_1(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44]), c_0_32])])).
cnf(c_0_50, plain, (v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_45, c_0_36])).
fof(c_0_51, plain, ![X35, X36, X37, X38]:((~r2_relset_1(X35,X36,X37,X38)|X37=X38|(~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(X35,X36)))|~m1_subset_1(X38,k1_zfmisc_1(k2_zfmisc_1(X35,X36)))))&(X37!=X38|r2_relset_1(X35,X36,X37,X38)|(~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(X35,X36)))|~m1_subset_1(X38,k1_zfmisc_1(k2_zfmisc_1(X35,X36)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r2_relset_1])])])).
cnf(c_0_52, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k5_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_53, plain, ![X39, X40, X41, X42, X43, X44]:((v1_funct_1(k1_partfun1(X39,X40,X41,X42,X43,X44))|(~v1_funct_1(X43)|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(X39,X40)))|(~v1_funct_1(X44)|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(X41,X42))))))&(m1_subset_1(k1_partfun1(X39,X40,X41,X42,X43,X44),k1_zfmisc_1(k2_zfmisc_1(X39,X42)))|(~v1_funct_1(X43)|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(X39,X40)))|(~v1_funct_1(X44)|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(X41,X42))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_partfun1])])])).
cnf(c_0_54, plain, (m1_subset_1(k2_funct_1(X1),k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v2_funct_1(X1)|k2_relset_1(X3,X2,X1)!=X2|~v1_funct_1(X1)|~v1_funct_2(X1,X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_55, negated_conjecture, (k2_relset_1(esk1_0,k2_relat_1(esk2_0),esk2_0)=k2_relat_1(esk2_0)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_56, negated_conjecture, (v2_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_57, negated_conjecture, (v1_funct_2(esk2_0,esk1_0,k2_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_43]), c_0_44]), c_0_32])])).
cnf(c_0_58, plain, (v1_funct_1(k2_funct_1(X1))|~v2_funct_1(X1)|k2_relset_1(X2,X3,X1)!=X3|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_59, plain, (X3=X4|~r2_relset_1(X1,X2,X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_60, negated_conjecture, (r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_61, negated_conjecture, (k1_partfun1(X1,X2,esk1_0,esk1_0,X3,esk2_0)=k5_relat_1(X3,esk2_0)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_31]), c_0_32])])).
cnf(c_0_62, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_63, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_64, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_65, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_66, negated_conjecture, (m1_subset_1(k2_funct_1(esk2_0),k1_zfmisc_1(k2_zfmisc_1(k2_relat_1(esk2_0),esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56]), c_0_49]), c_0_57]), c_0_32])])).
cnf(c_0_67, negated_conjecture, (v1_funct_1(k2_funct_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_55]), c_0_56]), c_0_49]), c_0_57]), c_0_32])])).
cnf(c_0_68, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)=esk2_0|~m1_subset_1(k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_31])])).
cnf(c_0_69, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)=k5_relat_1(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63])])).
cnf(c_0_70, negated_conjecture, (m1_subset_1(k1_partfun1(X1,X2,esk1_0,esk1_0,X3,esk2_0),k1_zfmisc_1(k2_zfmisc_1(X1,esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_31]), c_0_32])])).
cnf(c_0_71, negated_conjecture, (v1_funct_1(k1_partfun1(X1,X2,k2_relat_1(esk2_0),esk1_0,X3,k2_funct_1(esk2_0)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67])])).
cnf(c_0_72, negated_conjecture, (k1_partfun1(X1,X2,k2_relat_1(esk2_0),esk1_0,X3,k2_funct_1(esk2_0))=k5_relat_1(X3,k2_funct_1(esk2_0))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_66]), c_0_67])])).
fof(c_0_73, plain, ![X45]:((k5_relat_1(X45,k2_funct_1(X45))=k6_relat_1(k1_relat_1(X45))|~v2_funct_1(X45)|(~v1_relat_1(X45)|~v1_funct_1(X45)))&(k5_relat_1(k2_funct_1(X45),X45)=k6_relat_1(k2_relat_1(X45))|~v2_funct_1(X45)|(~v1_relat_1(X45)|~v1_funct_1(X45)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t61_funct_1])])])).
fof(c_0_74, plain, ![X46]:k6_partfun1(X46)=k6_relat_1(X46), inference(variable_rename,[status(thm)],[redefinition_k6_partfun1])).
fof(c_0_75, plain, ![X47, X48, X49, X50, X51, X52]:(~m1_subset_1(X51,k1_zfmisc_1(k2_zfmisc_1(X47,X48)))|~m1_subset_1(X52,k1_zfmisc_1(k2_zfmisc_1(X49,X50)))|k4_relset_1(X47,X48,X49,X50,X51,X52)=k5_relat_1(X51,X52)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_relset_1])])).
cnf(c_0_76, negated_conjecture, (k5_relat_1(esk3_0,esk2_0)=esk2_0|~m1_subset_1(k5_relat_1(esk3_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_69]), c_0_69])).
cnf(c_0_77, negated_conjecture, (m1_subset_1(k5_relat_1(esk3_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_62]), c_0_69]), c_0_63])])).
fof(c_0_78, plain, ![X53]:(v1_partfun1(k6_partfun1(X53),X53)&m1_subset_1(k6_partfun1(X53),k1_zfmisc_1(k2_zfmisc_1(X53,X53)))), inference(variable_rename,[status(thm)],[dt_k6_partfun1])).
cnf(c_0_79, negated_conjecture, (v1_funct_1(k1_partfun1(esk1_0,esk1_0,k2_relat_1(esk2_0),esk1_0,esk2_0,k2_funct_1(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_31]), c_0_32])])).
cnf(c_0_80, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,k2_relat_1(esk2_0),esk1_0,esk2_0,k2_funct_1(esk2_0))=k5_relat_1(esk2_0,k2_funct_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_31]), c_0_32])])).
cnf(c_0_81, plain, (k5_relat_1(X1,k2_funct_1(X1))=k6_relat_1(k1_relat_1(X1))|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_82, plain, (k6_partfun1(X1)=k6_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_74])).
fof(c_0_83, plain, ![X54, X55, X56]:((r2_relset_1(X54,X55,k4_relset_1(X54,X54,X54,X55,k6_partfun1(X54),X56),X56)|~m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(X54,X55))))&(r2_relset_1(X54,X55,k4_relset_1(X54,X55,X55,X55,X56,k6_partfun1(X55)),X56)|~m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(X54,X55))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_funct_2])])])).
cnf(c_0_84, plain, (k4_relset_1(X2,X3,X5,X6,X1,X4)=k5_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_75])).
fof(c_0_85, plain, ![X57, X58, X59]:((v1_funct_1(X59)|(~v1_funct_1(X59)|~v1_partfun1(X59,X57))|~m1_subset_1(X59,k1_zfmisc_1(k2_zfmisc_1(X57,X58))))&(v1_funct_2(X59,X57,X58)|(~v1_funct_1(X59)|~v1_partfun1(X59,X57))|~m1_subset_1(X59,k1_zfmisc_1(k2_zfmisc_1(X57,X58))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_2])])])).
fof(c_0_86, plain, ![X60, X61, X62, X63, X64, X65]:((~v2_funct_1(X62)|(~v1_funct_1(X64)|(~v1_funct_2(X64,X63,X60)|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(X63,X60))))|(~v1_funct_1(X65)|(~v1_funct_2(X65,X63,X60)|~m1_subset_1(X65,k1_zfmisc_1(k2_zfmisc_1(X63,X60))))|(~r2_relset_1(X63,X61,k1_partfun1(X63,X60,X60,X61,X64,X62),k1_partfun1(X63,X60,X60,X61,X65,X62))|r2_relset_1(X63,X60,X64,X65))))|X61=k1_xboole_0|X60=k1_xboole_0|(~v1_funct_1(X62)|(~v1_funct_2(X62,X60,X61)|~m1_subset_1(X62,k1_zfmisc_1(k2_zfmisc_1(X60,X61))))))&(((v1_funct_1(esk5_3(X60,X61,X62))|v2_funct_1(X62)|X61=k1_xboole_0|X60=k1_xboole_0|(~v1_funct_1(X62)|(~v1_funct_2(X62,X60,X61)|~m1_subset_1(X62,k1_zfmisc_1(k2_zfmisc_1(X60,X61))))))&((v1_funct_2(esk5_3(X60,X61,X62),esk4_3(X60,X61,X62),X60)|v2_funct_1(X62)|X61=k1_xboole_0|X60=k1_xboole_0|(~v1_funct_1(X62)|(~v1_funct_2(X62,X60,X61)|~m1_subset_1(X62,k1_zfmisc_1(k2_zfmisc_1(X60,X61))))))&(m1_subset_1(esk5_3(X60,X61,X62),k1_zfmisc_1(k2_zfmisc_1(esk4_3(X60,X61,X62),X60)))|v2_funct_1(X62)|X61=k1_xboole_0|X60=k1_xboole_0|(~v1_funct_1(X62)|(~v1_funct_2(X62,X60,X61)|~m1_subset_1(X62,k1_zfmisc_1(k2_zfmisc_1(X60,X61))))))))&(((v1_funct_1(esk6_3(X60,X61,X62))|v2_funct_1(X62)|X61=k1_xboole_0|X60=k1_xboole_0|(~v1_funct_1(X62)|(~v1_funct_2(X62,X60,X61)|~m1_subset_1(X62,k1_zfmisc_1(k2_zfmisc_1(X60,X61))))))&((v1_funct_2(esk6_3(X60,X61,X62),esk4_3(X60,X61,X62),X60)|v2_funct_1(X62)|X61=k1_xboole_0|X60=k1_xboole_0|(~v1_funct_1(X62)|(~v1_funct_2(X62,X60,X61)|~m1_subset_1(X62,k1_zfmisc_1(k2_zfmisc_1(X60,X61))))))&(m1_subset_1(esk6_3(X60,X61,X62),k1_zfmisc_1(k2_zfmisc_1(esk4_3(X60,X61,X62),X60)))|v2_funct_1(X62)|X61=k1_xboole_0|X60=k1_xboole_0|(~v1_funct_1(X62)|(~v1_funct_2(X62,X60,X61)|~m1_subset_1(X62,k1_zfmisc_1(k2_zfmisc_1(X60,X61))))))))&((r2_relset_1(esk4_3(X60,X61,X62),X61,k1_partfun1(esk4_3(X60,X61,X62),X60,X60,X61,esk5_3(X60,X61,X62),X62),k1_partfun1(esk4_3(X60,X61,X62),X60,X60,X61,esk6_3(X60,X61,X62),X62))|v2_funct_1(X62)|X61=k1_xboole_0|X60=k1_xboole_0|(~v1_funct_1(X62)|(~v1_funct_2(X62,X60,X61)|~m1_subset_1(X62,k1_zfmisc_1(k2_zfmisc_1(X60,X61))))))&(~r2_relset_1(esk4_3(X60,X61,X62),X60,esk5_3(X60,X61,X62),esk6_3(X60,X61,X62))|v2_funct_1(X62)|X61=k1_xboole_0|X60=k1_xboole_0|(~v1_funct_1(X62)|(~v1_funct_2(X62,X60,X61)|~m1_subset_1(X62,k1_zfmisc_1(k2_zfmisc_1(X60,X61)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t27_funct_2])])])])])])).
cnf(c_0_87, negated_conjecture, (k5_relat_1(esk3_0,esk2_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_77])])).
cnf(c_0_88, plain, (m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_89, negated_conjecture, (v1_funct_1(k5_relat_1(esk2_0,k2_funct_1(esk2_0)))), inference(rw,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_90, plain, (k5_relat_1(X1,k2_funct_1(X1))=k6_partfun1(k1_relat_1(X1))|~v1_funct_1(X1)|~v2_funct_1(X1)|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_91, plain, (r2_relset_1(X1,X2,k4_relset_1(X1,X1,X1,X2,k6_partfun1(X1),X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_92, negated_conjecture, (k4_relset_1(X1,X2,esk1_0,esk1_0,X3,esk2_0)=k5_relat_1(X3,esk2_0)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_84, c_0_31])).
cnf(c_0_93, plain, (v1_funct_2(X1,X2,X3)|~v1_funct_1(X1)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_94, plain, (v1_partfun1(k6_partfun1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_95, plain, (r2_relset_1(X3,X4,X2,X5)|X6=k1_xboole_0|X4=k1_xboole_0|~v2_funct_1(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X5)|~v1_funct_2(X5,X3,X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~r2_relset_1(X3,X6,k1_partfun1(X3,X4,X4,X6,X2,X1),k1_partfun1(X3,X4,X4,X6,X5,X1))|~v1_funct_1(X1)|~v1_funct_2(X1,X4,X6)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X6)))), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_96, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)=esk2_0), inference(rw,[status(thm)],[c_0_69, c_0_87])).
cnf(c_0_97, negated_conjecture, (v1_funct_2(esk3_0,esk1_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_98, negated_conjecture, (k1_partfun1(X1,X1,esk1_0,esk1_0,k6_partfun1(X1),esk2_0)=k5_relat_1(k6_partfun1(X1),esk2_0)|~v1_funct_1(k6_partfun1(X1))), inference(spm,[status(thm)],[c_0_61, c_0_88])).
cnf(c_0_99, negated_conjecture, (v1_funct_1(k6_partfun1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_43]), c_0_44]), c_0_56]), c_0_32])])).
cnf(c_0_100, negated_conjecture, (r2_relset_1(esk1_0,esk1_0,k4_relset_1(esk1_0,esk1_0,esk1_0,esk1_0,k6_partfun1(esk1_0),esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_91, c_0_31])).
cnf(c_0_101, negated_conjecture, (k4_relset_1(X1,X1,esk1_0,esk1_0,k6_partfun1(X1),esk2_0)=k5_relat_1(k6_partfun1(X1),esk2_0)), inference(spm,[status(thm)],[c_0_92, c_0_88])).
cnf(c_0_102, plain, (v1_funct_2(k6_partfun1(X1),X1,X1)|~v1_funct_1(k6_partfun1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_88]), c_0_94])])).
fof(c_0_103, plain, ![X69, X70]:((~m1_subset_1(X69,k1_zfmisc_1(X70))|r1_tarski(X69,X70))&(~r1_tarski(X69,X70)|m1_subset_1(X69,k1_zfmisc_1(X70)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_104, plain, ![X71, X72]:((k2_zfmisc_1(X71,X72)!=k1_xboole_0|(X71=k1_xboole_0|X72=k1_xboole_0))&((X71!=k1_xboole_0|k2_zfmisc_1(X71,X72)=k1_xboole_0)&(X72!=k1_xboole_0|k2_zfmisc_1(X71,X72)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_105, negated_conjecture, (esk1_0=k1_xboole_0|r2_relset_1(esk1_0,esk1_0,X1,esk3_0)|~r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))|~v1_funct_2(X1,esk1_0,esk1_0)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_56]), c_0_62]), c_0_31]), c_0_97]), c_0_30]), c_0_63]), c_0_32])])).
cnf(c_0_106, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,k6_partfun1(esk1_0),esk2_0)=k5_relat_1(k6_partfun1(esk1_0),esk2_0)), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_107, negated_conjecture, (r2_relset_1(esk1_0,esk1_0,k5_relat_1(k6_partfun1(esk1_0),esk2_0),esk2_0)), inference(rw,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_108, negated_conjecture, (v1_funct_2(k6_partfun1(esk1_0),esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_102, c_0_99])).
cnf(c_0_109, plain, (r2_relset_1(X3,X4,X1,X2)|X1!=X2|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_110, negated_conjecture, (k1_relset_1(esk1_0,esk1_0,esk3_0)=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_97]), c_0_62]), c_0_63])])).
cnf(c_0_111, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_112, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_113, negated_conjecture, (esk1_0=k1_xboole_0|r2_relset_1(esk1_0,esk1_0,k6_partfun1(esk1_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_107]), c_0_88]), c_0_108]), c_0_99])])).
cnf(c_0_114, plain, (r2_relset_1(X1,X2,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(er,[status(thm)],[c_0_109])).
cnf(c_0_115, negated_conjecture, (k1_relat_1(esk3_0)=esk1_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_62]), c_0_110])).
cnf(c_0_116, negated_conjecture, (v1_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_62]), c_0_40])])).
cnf(c_0_117, plain, (r1_tarski(k6_partfun1(X1),k2_zfmisc_1(X1,X1))), inference(spm,[status(thm)],[c_0_111, c_0_88])).
cnf(c_0_118, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_112])).
fof(c_0_119, plain, ![X73]:r1_tarski(k1_xboole_0,X73), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_120, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_121, negated_conjecture, (r1_tarski(esk2_0,k2_zfmisc_1(esk1_0,k2_relat_1(esk2_0)))), inference(spm,[status(thm)],[c_0_111, c_0_49])).
cnf(c_0_122, negated_conjecture, (~r2_relset_1(esk1_0,esk1_0,esk3_0,k6_partfun1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_123, negated_conjecture, (k6_partfun1(esk1_0)=esk3_0|esk1_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_113]), c_0_62]), c_0_88])])).
cnf(c_0_124, negated_conjecture, (r2_relset_1(esk1_0,esk1_0,esk3_0,esk3_0)), inference(spm,[status(thm)],[c_0_114, c_0_62])).
cnf(c_0_125, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_115]), c_0_116]), c_0_63])])).
cnf(c_0_126, plain, (r1_tarski(k6_partfun1(k1_xboole_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_117, c_0_118])).
cnf(c_0_127, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_128, negated_conjecture, (k2_zfmisc_1(esk1_0,k2_relat_1(esk2_0))=esk2_0|~r1_tarski(k2_zfmisc_1(esk1_0,k2_relat_1(esk2_0)),esk2_0)), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_129, negated_conjecture, (esk1_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_124])])).
cnf(c_0_130, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0)))), inference(spm,[status(thm)],[c_0_111, c_0_125])).
cnf(c_0_131, plain, (k6_partfun1(k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_126]), c_0_127])])).
cnf(c_0_132, negated_conjecture, (k1_xboole_0=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_128, c_0_129]), c_0_118]), c_0_129]), c_0_118]), c_0_127])])).
cnf(c_0_133, negated_conjecture, (k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0))=esk3_0|~r1_tarski(k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0)),esk3_0)), inference(spm,[status(thm)],[c_0_120, c_0_130])).
cnf(c_0_134, plain, (k6_partfun1(esk2_0)=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_131, c_0_132]), c_0_132])).
cnf(c_0_135, negated_conjecture, (esk2_0=esk3_0), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_133, c_0_129]), c_0_118]), c_0_129]), c_0_118]), c_0_127])]), c_0_132])).
cnf(c_0_136, plain, (r2_relset_1(X1,X1,k6_partfun1(X1),k6_partfun1(X1))), inference(spm,[status(thm)],[c_0_114, c_0_88])).
cnf(c_0_137, plain, (k6_partfun1(esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_134, c_0_135]), c_0_135])).
cnf(c_0_138, negated_conjecture, (~r2_relset_1(esk3_0,esk3_0,esk3_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_122, c_0_129]), c_0_129]), c_0_129]), c_0_131]), c_0_132]), c_0_135]), c_0_132]), c_0_135]), c_0_132]), c_0_135])).
cnf(c_0_139, plain, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_137]), c_0_138]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 140
# Proof object clause steps            : 95
# Proof object formula steps           : 45
# Proof object conjectures             : 57
# Proof object clause conjectures      : 54
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 36
# Proof object initial formulas used   : 22
# Proof object generating inferences   : 45
# Proof object simplifying inferences  : 116
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 53
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 50
# Processed clauses                    : 1855
# ...of these trivial                  : 97
# ...subsumed                          : 275
# ...remaining for further processing  : 1483
# Other redundant clauses eliminated   : 9
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 16
# Backward-rewritten                   : 1368
# Generated clauses                    : 9559
# ...of the previous two non-trivial   : 10367
# Contextual simplify-reflections      : 77
# Paramodulations                      : 9550
# Factorizations                       : 0
# Equation resolutions                 : 9
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 94
#    Positive orientable unit clauses  : 32
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 61
# Current number of unprocessed clauses: 8375
# ...number of literals in the above   : 18807
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1385
# Clause-clause subsumption calls (NU) : 23061
# Rec. Clause-clause subsumption calls : 10463
# Non-unit clause-clause subsumptions  : 368
# Unit Clause-clause subsumption calls : 7356
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 39744
# BW rewrite match successes           : 64
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 354053

# -------------------------------------------------
# User time                : 0.149 s
# System time              : 0.003 s
# Total time               : 0.153 s
# Maximum resident set size: 28716 pages
