# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t42_enumset1, axiom, ![X1, X2, X3]:k1_enumset1(X1,X2,X3)=k2_xboole_0(k1_tarski(X1),k2_tarski(X2,X3)), file('data/MPTP/proofs/0fe75b76818d4e93ba7b961eddd71184.E_input', t42_enumset1)).
fof(t41_enumset1, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_xboole_0(k1_tarski(X1),k1_tarski(X2)), file('data/MPTP/proofs/0fe75b76818d4e93ba7b961eddd71184.E_input', t41_enumset1)).
fof(t48_enumset1, axiom, ![X1, X2, X3, X4, X5]:k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_tarski(X1,X2),k1_enumset1(X3,X4,X5)), file('data/MPTP/proofs/0fe75b76818d4e93ba7b961eddd71184.E_input', t48_enumset1)).
fof(l68_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_enumset1(X5,X6,X7)), file('data/MPTP/proofs/0fe75b76818d4e93ba7b961eddd71184.E_input', l68_enumset1)).
fof(t51_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k1_tarski(X1),k3_enumset1(X2,X3,X4,X5,X6)), file('data/MPTP/proofs/0fe75b76818d4e93ba7b961eddd71184.E_input', t51_enumset1)).
fof(t57_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k2_tarski(X1,X2),k3_enumset1(X3,X4,X5,X6,X7)), file('data/MPTP/proofs/0fe75b76818d4e93ba7b961eddd71184.E_input', t57_enumset1)).
fof(t50_enumset1, axiom, ![X1, X2, X3, X4, X5]:k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_tarski(X5)), file('data/MPTP/proofs/0fe75b76818d4e93ba7b961eddd71184.E_input', t50_enumset1)).
fof(t62_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k1_tarski(X1),k5_enumset1(X2,X3,X4,X5,X6,X7,X8)), file('data/MPTP/proofs/0fe75b76818d4e93ba7b961eddd71184.E_input', t62_enumset1)).
fof(t56_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k1_tarski(X1),k4_enumset1(X2,X3,X4,X5,X6,X7)), file('data/MPTP/proofs/0fe75b76818d4e93ba7b961eddd71184.E_input', t56_enumset1)).
fof(t55_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_tarski(X6)), file('data/MPTP/proofs/0fe75b76818d4e93ba7b961eddd71184.E_input', t55_enumset1)).
fof(t6_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k2_xboole_0(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/0fe75b76818d4e93ba7b961eddd71184.E_input', t6_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/0fe75b76818d4e93ba7b961eddd71184.E_input', commutativity_k2_xboole_0)).
fof(t5_xboole_1, axiom, ![X1, X2, X3]:k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_xboole_0(X1,X3),k2_xboole_0(X2,X3)), file('data/MPTP/proofs/0fe75b76818d4e93ba7b961eddd71184.E_input', t5_xboole_1)).
fof(t68_enumset1, conjecture, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k5_enumset1(X1,X2,X3,X4,X5,X6,X7),k1_tarski(X8)), file('data/MPTP/proofs/0fe75b76818d4e93ba7b961eddd71184.E_input', t68_enumset1)).
fof(t65_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8)), file('data/MPTP/proofs/0fe75b76818d4e93ba7b961eddd71184.E_input', t65_enumset1)).
fof(idempotence_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/0fe75b76818d4e93ba7b961eddd71184.E_input', idempotence_k2_xboole_0)).
fof(c_0_16, plain, ![X78, X79, X80]:k1_enumset1(X78,X79,X80)=k2_xboole_0(k1_tarski(X78),k2_tarski(X79,X80)), inference(variable_rename,[status(thm)],[t42_enumset1])).
fof(c_0_17, plain, ![X31, X32]:k2_tarski(X31,X32)=k2_xboole_0(k1_tarski(X31),k1_tarski(X32)), inference(variable_rename,[status(thm)],[t41_enumset1])).
fof(c_0_18, plain, ![X17, X18, X19, X20, X21]:k3_enumset1(X17,X18,X19,X20,X21)=k2_xboole_0(k2_tarski(X17,X18),k1_enumset1(X19,X20,X21)), inference(variable_rename,[status(thm)],[t48_enumset1])).
cnf(c_0_19, plain, (k1_enumset1(X1,X2,X3)=k2_xboole_0(k1_tarski(X1),k2_tarski(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, (k2_tarski(X1,X2)=k2_xboole_0(k1_tarski(X1),k1_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_21, plain, ![X48, X49, X50, X51, X52, X53, X54]:k5_enumset1(X48,X49,X50,X51,X52,X53,X54)=k2_xboole_0(k2_enumset1(X48,X49,X50,X51),k1_enumset1(X52,X53,X54)), inference(variable_rename,[status(thm)],[l68_enumset1])).
fof(c_0_22, plain, ![X40, X41, X42, X43, X44, X45]:k4_enumset1(X40,X41,X42,X43,X44,X45)=k2_xboole_0(k1_tarski(X40),k3_enumset1(X41,X42,X43,X44,X45)), inference(variable_rename,[status(thm)],[t51_enumset1])).
cnf(c_0_23, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_tarski(X1,X2),k1_enumset1(X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (k1_enumset1(X1,X2,X3)=k2_xboole_0(k1_tarski(X1),k2_xboole_0(k1_tarski(X2),k1_tarski(X3)))), inference(rw,[status(thm)],[c_0_19, c_0_20])).
fof(c_0_25, plain, ![X55, X56, X57, X58, X59, X60, X61]:k5_enumset1(X55,X56,X57,X58,X59,X60,X61)=k2_xboole_0(k2_tarski(X55,X56),k3_enumset1(X57,X58,X59,X60,X61)), inference(variable_rename,[status(thm)],[t57_enumset1])).
cnf(c_0_26, plain, (k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_enumset1(X5,X6,X7))), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_27, plain, ![X62, X63, X64, X65, X66]:k3_enumset1(X62,X63,X64,X65,X66)=k2_xboole_0(k2_enumset1(X62,X63,X64,X65),k1_tarski(X66)), inference(variable_rename,[status(thm)],[t50_enumset1])).
fof(c_0_28, plain, ![X23, X24, X25, X26, X27, X28, X29, X30]:k6_enumset1(X23,X24,X25,X26,X27,X28,X29,X30)=k2_xboole_0(k1_tarski(X23),k5_enumset1(X24,X25,X26,X27,X28,X29,X30)), inference(variable_rename,[status(thm)],[t62_enumset1])).
fof(c_0_29, plain, ![X33, X34, X35, X36, X37, X38, X39]:k5_enumset1(X33,X34,X35,X36,X37,X38,X39)=k2_xboole_0(k1_tarski(X33),k4_enumset1(X34,X35,X36,X37,X38,X39)), inference(variable_rename,[status(thm)],[t56_enumset1])).
cnf(c_0_30, plain, (k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k1_tarski(X1),k3_enumset1(X2,X3,X4,X5,X6))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_xboole_0(k1_tarski(X1),k1_tarski(X2)),k2_xboole_0(k1_tarski(X3),k2_xboole_0(k1_tarski(X4),k1_tarski(X5))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_20]), c_0_24])).
cnf(c_0_32, plain, (k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k2_tarski(X1,X2),k3_enumset1(X3,X4,X5,X6,X7))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_xboole_0(k1_tarski(X5),k2_xboole_0(k1_tarski(X6),k1_tarski(X7))))), inference(rw,[status(thm)],[c_0_26, c_0_24])).
cnf(c_0_34, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_tarski(X5))), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_35, plain, ![X81, X82, X83, X84, X85, X86]:k4_enumset1(X81,X82,X83,X84,X85,X86)=k2_xboole_0(k3_enumset1(X81,X82,X83,X84,X85),k1_tarski(X86)), inference(variable_rename,[status(thm)],[t55_enumset1])).
fof(c_0_36, plain, ![X87, X88]:k2_xboole_0(X87,k2_xboole_0(X87,X88))=k2_xboole_0(X87,X88), inference(variable_rename,[status(thm)],[t6_xboole_1])).
fof(c_0_37, plain, ![X46, X47]:k2_xboole_0(X46,X47)=k2_xboole_0(X47,X46), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_38, plain, ![X67, X68, X69]:k2_xboole_0(k2_xboole_0(X67,X68),X69)=k2_xboole_0(k2_xboole_0(X67,X69),k2_xboole_0(X68,X69)), inference(variable_rename,[status(thm)],[t5_xboole_1])).
fof(c_0_39, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k5_enumset1(X1,X2,X3,X4,X5,X6,X7),k1_tarski(X8))), inference(assume_negation,[status(cth)],[t68_enumset1])).
fof(c_0_40, plain, ![X70, X71, X72, X73, X74, X75, X76, X77]:k6_enumset1(X70,X71,X72,X73,X74,X75,X76,X77)=k2_xboole_0(k2_enumset1(X70,X71,X72,X73),k2_enumset1(X74,X75,X76,X77)), inference(variable_rename,[status(thm)],[t65_enumset1])).
cnf(c_0_41, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k1_tarski(X1),k5_enumset1(X2,X3,X4,X5,X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_42, plain, (k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k1_tarski(X1),k4_enumset1(X2,X3,X4,X5,X6,X7))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_43, plain, (k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k1_tarski(X1),k2_xboole_0(k2_xboole_0(k1_tarski(X2),k1_tarski(X3)),k2_xboole_0(k1_tarski(X4),k2_xboole_0(k1_tarski(X5),k1_tarski(X6)))))), inference(rw,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_44, plain, (k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_xboole_0(k1_tarski(X5),k2_xboole_0(k1_tarski(X6),k1_tarski(X7))))=k2_xboole_0(k2_xboole_0(k1_tarski(X1),k1_tarski(X2)),k2_xboole_0(k2_xboole_0(k1_tarski(X3),k1_tarski(X4)),k2_xboole_0(k1_tarski(X5),k2_xboole_0(k1_tarski(X6),k1_tarski(X7)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_20]), c_0_31]), c_0_33])).
cnf(c_0_45, plain, (k2_xboole_0(k2_xboole_0(k1_tarski(X1),k1_tarski(X2)),k2_xboole_0(k1_tarski(X3),k2_xboole_0(k1_tarski(X4),k1_tarski(X5))))=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_tarski(X5))), inference(rw,[status(thm)],[c_0_34, c_0_31])).
cnf(c_0_46, plain, (k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_tarski(X6))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_47, plain, (k2_xboole_0(X1,k2_xboole_0(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_48, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_49, plain, (k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_xboole_0(X1,X3),k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_50, negated_conjecture, k6_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0)!=k2_xboole_0(k5_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0),k1_tarski(esk8_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_39])])])).
cnf(c_0_51, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_52, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k1_tarski(X1),k2_xboole_0(k2_enumset1(X2,X3,X4,X5),k2_xboole_0(k1_tarski(X6),k2_xboole_0(k1_tarski(X7),k1_tarski(X8)))))), inference(rw,[status(thm)],[c_0_41, c_0_33])).
cnf(c_0_53, plain, (k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_xboole_0(k1_tarski(X5),k2_xboole_0(k1_tarski(X6),k1_tarski(X7))))=k2_xboole_0(k1_tarski(X1),k2_xboole_0(k1_tarski(X2),k2_xboole_0(k2_xboole_0(k1_tarski(X3),k1_tarski(X4)),k2_xboole_0(k1_tarski(X5),k2_xboole_0(k1_tarski(X6),k1_tarski(X7))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_43]), c_0_33])).
cnf(c_0_54, plain, (k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_xboole_0(k1_tarski(X5),k2_xboole_0(k1_tarski(X6),k1_tarski(X7))))=k2_xboole_0(k2_xboole_0(k1_tarski(X1),k1_tarski(X2)),k2_xboole_0(k2_enumset1(X3,X4,X5,X6),k1_tarski(X7)))), inference(rw,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_55, plain, (k2_xboole_0(k1_tarski(X1),k2_xboole_0(k2_xboole_0(k1_tarski(X2),k1_tarski(X3)),k2_xboole_0(k1_tarski(X4),k2_xboole_0(k1_tarski(X5),k1_tarski(X6)))))=k2_xboole_0(k2_xboole_0(k2_xboole_0(k1_tarski(X1),k1_tarski(X2)),k2_xboole_0(k1_tarski(X3),k2_xboole_0(k1_tarski(X4),k1_tarski(X5)))),k1_tarski(X6))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_31]), c_0_43])).
cnf(c_0_56, plain, (k2_xboole_0(X1,k2_xboole_0(X2,X1))=k2_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_57, plain, (k2_xboole_0(k2_xboole_0(X1,X2),k2_xboole_0(X3,X1))=k2_xboole_0(k2_xboole_0(X2,X3),X1)), inference(spm,[status(thm)],[c_0_49, c_0_48])).
cnf(c_0_58, negated_conjecture, (k6_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0)!=k2_xboole_0(k5_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0),k1_tarski(esk8_0))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_59, plain, (k2_xboole_0(k1_tarski(X1),k2_xboole_0(k2_enumset1(X2,X3,X4,X5),k2_xboole_0(k1_tarski(X6),k2_xboole_0(k1_tarski(X7),k1_tarski(X8)))))=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8))), inference(rw,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_60, plain, (k2_xboole_0(k2_xboole_0(k1_tarski(X1),k1_tarski(X2)),k2_xboole_0(k2_enumset1(X3,X4,X5,X6),k1_tarski(X7)))=k2_xboole_0(k1_tarski(X1),k2_xboole_0(k1_tarski(X2),k2_xboole_0(k2_enumset1(X3,X4,X5,X6),k1_tarski(X7))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_45]), c_0_54])).
cnf(c_0_61, plain, (k2_xboole_0(k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_tarski(X5)),k1_tarski(X6))=k2_xboole_0(k1_tarski(X1),k2_xboole_0(k2_enumset1(X2,X3,X4,X5),k1_tarski(X6)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_45]), c_0_45])).
cnf(c_0_62, plain, (k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_56]), c_0_56]), c_0_57])).
cnf(c_0_63, negated_conjecture, (k2_xboole_0(k1_tarski(esk1_0),k2_xboole_0(k2_enumset1(esk2_0,esk3_0,esk4_0,esk5_0),k2_xboole_0(k1_tarski(esk6_0),k2_xboole_0(k1_tarski(esk7_0),k1_tarski(esk8_0)))))!=k2_xboole_0(k2_xboole_0(k2_enumset1(esk1_0,esk2_0,esk3_0,esk4_0),k2_xboole_0(k1_tarski(esk5_0),k2_xboole_0(k1_tarski(esk6_0),k1_tarski(esk7_0)))),k1_tarski(esk8_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_33]), c_0_52])).
cnf(c_0_64, plain, (k2_xboole_0(k1_tarski(X1),k2_xboole_0(k1_tarski(X2),k2_xboole_0(k1_tarski(X3),k2_xboole_0(k2_enumset1(X4,X5,X6,X7),k1_tarski(X8)))))=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_54]), c_0_60])).
fof(c_0_65, plain, ![X22]:k2_xboole_0(X22,X22)=X22, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k2_xboole_0])])).
cnf(c_0_66, plain, (k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_xboole_0(k1_tarski(X5),k1_tarski(X6)))=k2_xboole_0(k1_tarski(X1),k2_xboole_0(k2_enumset1(X2,X3,X4,X5),k1_tarski(X6)))), inference(rw,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_67, negated_conjecture, (k2_xboole_0(k1_tarski(esk8_0),k2_xboole_0(k2_xboole_0(k1_tarski(esk5_0),k2_xboole_0(k1_tarski(esk6_0),k1_tarski(esk7_0))),k2_enumset1(esk1_0,esk2_0,esk3_0,esk4_0)))!=k2_xboole_0(k1_tarski(esk1_0),k2_xboole_0(k2_xboole_0(k1_tarski(esk6_0),k2_xboole_0(k1_tarski(esk7_0),k1_tarski(esk8_0))),k2_enumset1(esk2_0,esk3_0,esk4_0,esk5_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_48]), c_0_48]), c_0_48])).
cnf(c_0_68, plain, (k2_xboole_0(k1_tarski(X1),k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8)))=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8))), inference(spm,[status(thm)],[c_0_47, c_0_64])).
cnf(c_0_69, plain, (k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_70, plain, (k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8))=k2_xboole_0(k2_enumset1(X8,X1,X2,X3),k2_enumset1(X4,X5,X6,X7))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_64]), c_0_62]), c_0_62]), c_0_62]), c_0_66]), c_0_64])).
cnf(c_0_71, negated_conjecture, (k2_xboole_0(k1_tarski(esk8_0),k2_xboole_0(k1_tarski(esk5_0),k2_xboole_0(k1_tarski(esk6_0),k2_xboole_0(k1_tarski(esk7_0),k2_enumset1(esk1_0,esk2_0,esk3_0,esk4_0)))))!=k2_xboole_0(k1_tarski(esk1_0),k2_xboole_0(k1_tarski(esk6_0),k2_xboole_0(k1_tarski(esk7_0),k2_xboole_0(k1_tarski(esk8_0),k2_enumset1(esk2_0,esk3_0,esk4_0,esk5_0)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_62]), c_0_62]), c_0_62]), c_0_62])).
cnf(c_0_72, plain, (k2_xboole_0(X1,k2_xboole_0(X2,X3))=k2_xboole_0(X3,k2_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_48, c_0_62])).
cnf(c_0_73, plain, (k2_xboole_0(k1_tarski(X1),k2_enumset1(X1,X2,X3,X4))=k2_enumset1(X1,X2,X3,X4)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_74, plain, (k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X1,X2,X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_69])).
cnf(c_0_75, negated_conjecture, (k2_xboole_0(k1_tarski(esk5_0),k2_xboole_0(k1_tarski(esk6_0),k2_xboole_0(k1_tarski(esk7_0),k2_xboole_0(k1_tarski(esk8_0),k2_enumset1(esk1_0,esk2_0,esk3_0,esk4_0)))))!=k2_xboole_0(k1_tarski(esk1_0),k2_xboole_0(k1_tarski(esk6_0),k2_xboole_0(k1_tarski(esk7_0),k2_xboole_0(k1_tarski(esk8_0),k2_enumset1(esk2_0,esk3_0,esk4_0,esk5_0)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_72]), c_0_62]), c_0_62]), c_0_48])).
cnf(c_0_76, plain, (k2_xboole_0(k1_tarski(X1),k2_xboole_0(k1_tarski(X2),k2_xboole_0(k1_tarski(X3),k2_xboole_0(k1_tarski(X4),k2_enumset1(X5,X6,X7,X8)))))=k2_xboole_0(k2_enumset1(X1,X2,X3,X5),k2_enumset1(X6,X7,X8,X4))), inference(spm,[status(thm)],[c_0_64, c_0_48])).
cnf(c_0_77, plain, (k2_xboole_0(k1_tarski(X1),k2_xboole_0(k1_tarski(X2),k2_xboole_0(k1_tarski(X3),k2_xboole_0(k1_tarski(X4),k1_tarski(X5)))))=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_tarski(X5))), inference(rw,[status(thm)],[c_0_45, c_0_62])).
cnf(c_0_78, plain, (k2_xboole_0(k1_tarski(X1),k2_enumset1(X2,X3,X4,X1))=k2_enumset1(X2,X3,X4,X1)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_79, negated_conjecture, (k2_xboole_0(k2_enumset1(esk1_0,esk6_0,esk7_0,esk2_0),k2_enumset1(esk3_0,esk4_0,esk5_0,esk8_0))!=k2_xboole_0(k2_enumset1(esk1_0,esk5_0,esk6_0,esk7_0),k2_enumset1(esk2_0,esk3_0,esk4_0,esk8_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_76]), c_0_74]), c_0_76])).
cnf(c_0_80, plain, (k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8))=k2_xboole_0(k2_enumset1(X1,X2,X8,X3),k2_enumset1(X4,X5,X6,X7))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_48]), c_0_61]), c_0_64])).
cnf(c_0_81, plain, (k2_xboole_0(k1_tarski(X1),k2_xboole_0(k1_tarski(X2),k2_xboole_0(k1_tarski(X3),k1_tarski(X4))))=k2_enumset1(X1,X2,X3,X4)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_69]), c_0_48]), c_0_78])).
cnf(c_0_82, negated_conjecture, (k2_xboole_0(k2_enumset1(esk1_0,esk6_0,esk2_0,esk3_0),k2_enumset1(esk4_0,esk5_0,esk8_0,esk7_0))!=k2_xboole_0(k2_enumset1(esk1_0,esk5_0,esk6_0,esk7_0),k2_enumset1(esk2_0,esk3_0,esk4_0,esk8_0))), inference(rw,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_83, plain, (k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X2,X4,X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_48]), c_0_81])).
cnf(c_0_84, negated_conjecture, (k2_xboole_0(k2_enumset1(esk1_0,esk6_0,esk2_0,esk3_0),k2_enumset1(esk4_0,esk5_0,esk7_0,esk8_0))!=k2_xboole_0(k2_enumset1(esk1_0,esk5_0,esk6_0,esk7_0),k2_enumset1(esk2_0,esk3_0,esk4_0,esk8_0))), inference(rw,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_85, plain, (k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X4,X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_48]), c_0_62]), c_0_81])).
cnf(c_0_86, plain, (k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_tarski(X5))=k2_xboole_0(k1_tarski(X1),k2_enumset1(X2,X3,X4,X5))), inference(rw,[status(thm)],[c_0_77, c_0_81])).
cnf(c_0_87, plain, (k2_xboole_0(X1,k2_xboole_0(X2,X3))=k2_xboole_0(X2,k2_xboole_0(X1,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_48]), c_0_62])).
cnf(c_0_88, negated_conjecture, (k2_xboole_0(k2_enumset1(esk1_0,esk5_0,esk6_0,esk7_0),k2_enumset1(esk2_0,esk3_0,esk4_0,esk8_0))!=k2_xboole_0(k2_enumset1(esk1_0,esk2_0,esk3_0,esk6_0),k2_enumset1(esk4_0,esk5_0,esk7_0,esk8_0))), inference(rw,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_89, plain, (k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8))=k2_xboole_0(k2_enumset1(X1,X2,X3,X8),k2_enumset1(X4,X5,X6,X7))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_86]), c_0_76])).
cnf(c_0_90, plain, (k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X3,X2,X4)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_87]), c_0_81])).
cnf(c_0_91, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_88, c_0_89]), c_0_83]), c_0_90]), c_0_83]), c_0_89]), c_0_83]), c_0_83]), c_0_90]), c_0_89]), c_0_74]), c_0_89]), c_0_83]), c_0_90])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 92
# Proof object clause steps            : 59
# Proof object formula steps           : 33
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 77
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 6
# Initial clauses in saturation        : 10
# Processed clauses                    : 226
# ...of these trivial                  : 77
# ...subsumed                          : 85
# ...remaining for further processing  : 64
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 24
# Generated clauses                    : 2089
# ...of the previous two non-trivial   : 1584
# Contextual simplify-reflections      : 0
# Paramodulations                      : 2089
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 40
#    Positive orientable unit clauses  : 24
#    Positive unorientable unit clauses: 16
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 1286
# ...number of literals in the above   : 1286
# Current number of archived formulas  : 0
# Current number of archived clauses   : 30
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 149
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 872
# BW rewrite match successes           : 612
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 21656

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 28656 pages
