# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t38_waybel_7, conjecture, ![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v5_waybel_6(X2,X1)=>v4_waybel_7(X2,X1)))), file('data/MPTP/proofs/10219f468f12489c88ab4cf1f3de4265.E_input', t38_waybel_7)).
fof(d6_waybel_7, axiom, ![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v4_waybel_7(X2,X1)<=>?[X3]:((~(v1_xboole_0(X3))&(v1_waybel_0(X3,X1)&(v12_waybel_0(X3,X1)&(v1_waybel_7(X3,X1)&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))))&X2=k1_yellow_0(X1,X3))))), file('data/MPTP/proofs/10219f468f12489c88ab4cf1f3de4265.E_input', d6_waybel_7)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/10219f468f12489c88ab4cf1f3de4265.E_input', dt_k1_yellow_0)).
fof(t37_waybel_7, axiom, ![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v5_waybel_6(X2,X1)=>v1_waybel_7(k5_waybel_0(X1,X2),X1)))), file('data/MPTP/proofs/10219f468f12489c88ab4cf1f3de4265.E_input', t37_waybel_7)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/10219f468f12489c88ab4cf1f3de4265.E_input', d9_lattice3)).
fof(t30_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4)))))&((r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))=>(X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3)))))), file('data/MPTP/proofs/10219f468f12489c88ab4cf1f3de4265.E_input', t30_yellow_0)).
fof(reflexivity_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>r3_orders_2(X1,X2,X2)), file('data/MPTP/proofs/10219f468f12489c88ab4cf1f3de4265.E_input', reflexivity_r3_orders_2)).
fof(fc12_waybel_0, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))&m1_subset_1(X2,u1_struct_0(X1)))=>v12_waybel_0(k5_waybel_0(X1,X2),X1)), file('data/MPTP/proofs/10219f468f12489c88ab4cf1f3de4265.E_input', fc12_waybel_0)).
fof(t17_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2))))), file('data/MPTP/proofs/10219f468f12489c88ab4cf1f3de4265.E_input', t17_waybel_0)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/10219f468f12489c88ab4cf1f3de4265.E_input', redefinition_r3_orders_2)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/10219f468f12489c88ab4cf1f3de4265.E_input', t4_subset)).
fof(dt_k5_waybel_0, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/10219f468f12489c88ab4cf1f3de4265.E_input', dt_k5_waybel_0)).
fof(fc8_waybel_0, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&m1_subset_1(X2,u1_struct_0(X1)))=>(~(v1_xboole_0(k5_waybel_0(X1,X2)))&v1_waybel_0(k5_waybel_0(X1,X2),X1))), file('data/MPTP/proofs/10219f468f12489c88ab4cf1f3de4265.E_input', fc8_waybel_0)).
fof(cc1_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v1_lattice3(X1)=>~(v2_struct_0(X1)))), file('data/MPTP/proofs/10219f468f12489c88ab4cf1f3de4265.E_input', cc1_lattice3)).
fof(c_0_14, negated_conjecture, ~(![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v5_waybel_6(X2,X1)=>v4_waybel_7(X2,X1))))), inference(assume_negation,[status(cth)],[t38_waybel_7])).
fof(c_0_15, plain, ![X38, X39, X41]:((((~v1_xboole_0(esk5_2(X38,X39))|~v4_waybel_7(X39,X38)|~m1_subset_1(X39,u1_struct_0(X38))|(~v3_orders_2(X38)|(~v4_orders_2(X38)|(~v5_orders_2(X38)|(~v1_lattice3(X38)|(~v2_lattice3(X38)|~l1_orders_2(X38)))))))&((v1_waybel_0(esk5_2(X38,X39),X38)|~v4_waybel_7(X39,X38)|~m1_subset_1(X39,u1_struct_0(X38))|(~v3_orders_2(X38)|(~v4_orders_2(X38)|(~v5_orders_2(X38)|(~v1_lattice3(X38)|(~v2_lattice3(X38)|~l1_orders_2(X38)))))))&((v12_waybel_0(esk5_2(X38,X39),X38)|~v4_waybel_7(X39,X38)|~m1_subset_1(X39,u1_struct_0(X38))|(~v3_orders_2(X38)|(~v4_orders_2(X38)|(~v5_orders_2(X38)|(~v1_lattice3(X38)|(~v2_lattice3(X38)|~l1_orders_2(X38)))))))&((v1_waybel_7(esk5_2(X38,X39),X38)|~v4_waybel_7(X39,X38)|~m1_subset_1(X39,u1_struct_0(X38))|(~v3_orders_2(X38)|(~v4_orders_2(X38)|(~v5_orders_2(X38)|(~v1_lattice3(X38)|(~v2_lattice3(X38)|~l1_orders_2(X38)))))))&(m1_subset_1(esk5_2(X38,X39),k1_zfmisc_1(u1_struct_0(X38)))|~v4_waybel_7(X39,X38)|~m1_subset_1(X39,u1_struct_0(X38))|(~v3_orders_2(X38)|(~v4_orders_2(X38)|(~v5_orders_2(X38)|(~v1_lattice3(X38)|(~v2_lattice3(X38)|~l1_orders_2(X38)))))))))))&(X39=k1_yellow_0(X38,esk5_2(X38,X39))|~v4_waybel_7(X39,X38)|~m1_subset_1(X39,u1_struct_0(X38))|(~v3_orders_2(X38)|(~v4_orders_2(X38)|(~v5_orders_2(X38)|(~v1_lattice3(X38)|(~v2_lattice3(X38)|~l1_orders_2(X38))))))))&(v1_xboole_0(X41)|(~v1_waybel_0(X41,X38)|(~v12_waybel_0(X41,X38)|(~v1_waybel_7(X41,X38)|~m1_subset_1(X41,k1_zfmisc_1(u1_struct_0(X38))))))|X39!=k1_yellow_0(X38,X41)|v4_waybel_7(X39,X38)|~m1_subset_1(X39,u1_struct_0(X38))|(~v3_orders_2(X38)|(~v4_orders_2(X38)|(~v5_orders_2(X38)|(~v1_lattice3(X38)|(~v2_lattice3(X38)|~l1_orders_2(X38)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_waybel_7])])])])])])).
fof(c_0_16, plain, ![X28, X29]:(~l1_orders_2(X28)|m1_subset_1(k1_yellow_0(X28,X29),u1_struct_0(X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
fof(c_0_17, plain, ![X26, X27]:(~v3_orders_2(X26)|(~v4_orders_2(X26)|(~v5_orders_2(X26)|(~v1_lattice3(X26)|(~v2_lattice3(X26)|~l1_orders_2(X26)))))|(~m1_subset_1(X27,u1_struct_0(X26))|(~v5_waybel_6(X27,X26)|v1_waybel_7(k5_waybel_0(X26,X27),X26)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t37_waybel_7])])])).
fof(c_0_18, negated_conjecture, ((v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_lattice3(esk1_0)&(v2_lattice3(esk1_0)&l1_orders_2(esk1_0))))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(v5_waybel_6(esk2_0,esk1_0)&~v4_waybel_7(esk2_0,esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
cnf(c_0_19, plain, (v1_xboole_0(X1)|v4_waybel_7(X3,X2)|~v1_waybel_0(X1,X2)|~v12_waybel_0(X1,X2)|~v1_waybel_7(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|X3!=k1_yellow_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v2_lattice3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, plain, (v1_waybel_7(k5_waybel_0(X1,X2),X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v5_waybel_6(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, negated_conjecture, (v5_waybel_6(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, negated_conjecture, (v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, negated_conjecture, (v1_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_29, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_30, plain, ![X12, X13, X14, X15]:((~r2_lattice3(X12,X13,X14)|(~m1_subset_1(X15,u1_struct_0(X12))|(~r2_hidden(X15,X13)|r1_orders_2(X12,X15,X14)))|~m1_subset_1(X14,u1_struct_0(X12))|~l1_orders_2(X12))&((m1_subset_1(esk3_3(X12,X13,X14),u1_struct_0(X12))|r2_lattice3(X12,X13,X14)|~m1_subset_1(X14,u1_struct_0(X12))|~l1_orders_2(X12))&((r2_hidden(esk3_3(X12,X13,X14),X13)|r2_lattice3(X12,X13,X14)|~m1_subset_1(X14,u1_struct_0(X12))|~l1_orders_2(X12))&(~r1_orders_2(X12,esk3_3(X12,X13,X14),X14)|r2_lattice3(X12,X13,X14)|~m1_subset_1(X14,u1_struct_0(X12))|~l1_orders_2(X12))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
fof(c_0_31, plain, ![X17, X18, X19, X20, X21]:(((r2_lattice3(X17,X19,X18)|(X18!=k1_yellow_0(X17,X19)|~r1_yellow_0(X17,X19))|~m1_subset_1(X18,u1_struct_0(X17))|(~v5_orders_2(X17)|~l1_orders_2(X17)))&(~m1_subset_1(X20,u1_struct_0(X17))|(~r2_lattice3(X17,X19,X20)|r1_orders_2(X17,X18,X20))|(X18!=k1_yellow_0(X17,X19)|~r1_yellow_0(X17,X19))|~m1_subset_1(X18,u1_struct_0(X17))|(~v5_orders_2(X17)|~l1_orders_2(X17))))&(((X18=k1_yellow_0(X17,X21)|(m1_subset_1(esk4_3(X17,X18,X21),u1_struct_0(X17))|~r2_lattice3(X17,X21,X18))|~m1_subset_1(X18,u1_struct_0(X17))|(~v5_orders_2(X17)|~l1_orders_2(X17)))&(r1_yellow_0(X17,X21)|(m1_subset_1(esk4_3(X17,X18,X21),u1_struct_0(X17))|~r2_lattice3(X17,X21,X18))|~m1_subset_1(X18,u1_struct_0(X17))|(~v5_orders_2(X17)|~l1_orders_2(X17))))&(((X18=k1_yellow_0(X17,X21)|(r2_lattice3(X17,X21,esk4_3(X17,X18,X21))|~r2_lattice3(X17,X21,X18))|~m1_subset_1(X18,u1_struct_0(X17))|(~v5_orders_2(X17)|~l1_orders_2(X17)))&(r1_yellow_0(X17,X21)|(r2_lattice3(X17,X21,esk4_3(X17,X18,X21))|~r2_lattice3(X17,X21,X18))|~m1_subset_1(X18,u1_struct_0(X17))|(~v5_orders_2(X17)|~l1_orders_2(X17))))&((X18=k1_yellow_0(X17,X21)|(~r1_orders_2(X17,X18,esk4_3(X17,X18,X21))|~r2_lattice3(X17,X21,X18))|~m1_subset_1(X18,u1_struct_0(X17))|(~v5_orders_2(X17)|~l1_orders_2(X17)))&(r1_yellow_0(X17,X21)|(~r1_orders_2(X17,X18,esk4_3(X17,X18,X21))|~r2_lattice3(X17,X21,X18))|~m1_subset_1(X18,u1_struct_0(X17))|(~v5_orders_2(X17)|~l1_orders_2(X17))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_yellow_0])])])])])])).
fof(c_0_32, plain, ![X9, X10, X11]:(v2_struct_0(X9)|(~v3_orders_2(X9)|~l1_orders_2(X9))|(~m1_subset_1(X10,u1_struct_0(X9))|~m1_subset_1(X11,u1_struct_0(X9)))|r3_orders_2(X9,X10,X10)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r3_orders_2])])])).
cnf(c_0_33, plain, (v1_xboole_0(X1)|v4_waybel_7(k1_yellow_0(X2,X1),X2)|~v12_waybel_0(X1,X2)|~v1_waybel_7(X1,X2)|~v1_waybel_0(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_19]), c_0_20])).
cnf(c_0_34, negated_conjecture, (v1_waybel_7(k5_waybel_0(esk1_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29])])).
fof(c_0_35, plain, ![X42, X43]:(v2_struct_0(X42)|(~v4_orders_2(X42)|~l1_orders_2(X42))|~m1_subset_1(X43,u1_struct_0(X42))|v12_waybel_0(k5_waybel_0(X42,X43),X42)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc12_waybel_0])])])).
cnf(c_0_36, plain, (r1_orders_2(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37, plain, (X1=k1_yellow_0(X2,X3)|r2_lattice3(X2,X3,esk4_3(X2,X1,X3))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, plain, (X1=k1_yellow_0(X2,X3)|m1_subset_1(esk4_3(X2,X1,X3),u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_39, plain, ![X23, X24, X25]:((~r2_hidden(X25,k5_waybel_0(X23,X24))|r1_orders_2(X23,X25,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~m1_subset_1(X24,u1_struct_0(X23))|(v2_struct_0(X23)|~l1_orders_2(X23)))&(~r1_orders_2(X23,X25,X24)|r2_hidden(X25,k5_waybel_0(X23,X24))|~m1_subset_1(X25,u1_struct_0(X23))|~m1_subset_1(X24,u1_struct_0(X23))|(v2_struct_0(X23)|~l1_orders_2(X23)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_waybel_0])])])])])).
fof(c_0_40, plain, ![X32, X33, X34]:((~r3_orders_2(X32,X33,X34)|r1_orders_2(X32,X33,X34)|(v2_struct_0(X32)|(~v3_orders_2(X32)|~l1_orders_2(X32))|(~m1_subset_1(X33,u1_struct_0(X32))|~m1_subset_1(X34,u1_struct_0(X32)))))&(~r1_orders_2(X32,X33,X34)|r3_orders_2(X32,X33,X34)|(v2_struct_0(X32)|(~v3_orders_2(X32)|~l1_orders_2(X32))|(~m1_subset_1(X33,u1_struct_0(X32))|~m1_subset_1(X34,u1_struct_0(X32)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_41, plain, (v2_struct_0(X1)|r3_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_42, plain, ![X35, X36, X37]:(~r2_hidden(X35,X36)|~m1_subset_1(X36,k1_zfmisc_1(X37))|m1_subset_1(X35,X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_43, plain, ![X30, X31]:(v2_struct_0(X30)|~l1_orders_2(X30)|~m1_subset_1(X31,u1_struct_0(X30))|m1_subset_1(k5_waybel_0(X30,X31),k1_zfmisc_1(u1_struct_0(X30)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_waybel_0])])])).
cnf(c_0_44, negated_conjecture, (v1_xboole_0(k5_waybel_0(esk1_0,esk2_0))|v4_waybel_7(k1_yellow_0(esk1_0,k5_waybel_0(esk1_0,esk2_0)),esk1_0)|~v12_waybel_0(k5_waybel_0(esk1_0,esk2_0),esk1_0)|~v1_waybel_0(k5_waybel_0(esk1_0,esk2_0),esk1_0)|~m1_subset_1(k5_waybel_0(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29])])).
cnf(c_0_45, plain, (v2_struct_0(X1)|v12_waybel_0(k5_waybel_0(X1,X2),X1)|~v4_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_46, plain, ![X7, X8]:((~v1_xboole_0(k5_waybel_0(X7,X8))|(v2_struct_0(X7)|(~v3_orders_2(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,u1_struct_0(X7))))&(v1_waybel_0(k5_waybel_0(X7,X8),X7)|(v2_struct_0(X7)|(~v3_orders_2(X7)|~l1_orders_2(X7))|~m1_subset_1(X8,u1_struct_0(X7))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc8_waybel_0])])])])).
cnf(c_0_47, plain, (X1=k1_yellow_0(X2,X3)|r1_orders_2(X2,X4,esk4_3(X2,X1,X3))|~r2_hidden(X4,X3)|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_48, plain, (r2_hidden(X2,k5_waybel_0(X1,X3))|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_49, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_50, negated_conjecture, (r3_orders_2(esk1_0,X1,X1)|v2_struct_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_23]), c_0_24]), c_0_29])])).
cnf(c_0_51, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_52, plain, (v2_struct_0(X1)|m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_53, negated_conjecture, (v1_xboole_0(k5_waybel_0(esk1_0,esk2_0))|v2_struct_0(esk1_0)|v4_waybel_7(k1_yellow_0(esk1_0,k5_waybel_0(esk1_0,esk2_0)),esk1_0)|~v1_waybel_0(k5_waybel_0(esk1_0,esk2_0),esk1_0)|~m1_subset_1(k5_waybel_0(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_23]), c_0_24]), c_0_28])])).
cnf(c_0_54, plain, (v1_waybel_0(k5_waybel_0(X1,X2),X1)|v2_struct_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_55, plain, (X1=k1_yellow_0(X2,k5_waybel_0(X3,X4))|r1_orders_2(X2,X5,esk4_3(X2,X1,k5_waybel_0(X3,X4)))|v2_struct_0(X3)|~r1_orders_2(X3,X5,X4)|~r2_lattice3(X2,k5_waybel_0(X3,X4),X1)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X3))|~m1_subset_1(X5,u1_struct_0(X3))|~l1_orders_2(X2)|~l1_orders_2(X3)|~v5_orders_2(X2)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_56, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|v2_struct_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_24]), c_0_29])])).
cnf(c_0_57, plain, (r1_orders_2(X2,X1,X3)|v2_struct_0(X2)|~r2_hidden(X1,k5_waybel_0(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_58, plain, (v2_struct_0(X1)|m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X2,k5_waybel_0(X1,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_59, plain, (v2_struct_0(X1)|~v1_xboole_0(k5_waybel_0(X1,X2))|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_60, negated_conjecture, (v1_xboole_0(k5_waybel_0(esk1_0,esk2_0))|v2_struct_0(esk1_0)|v4_waybel_7(k1_yellow_0(esk1_0,k5_waybel_0(esk1_0,esk2_0)),esk1_0)|~m1_subset_1(k5_waybel_0(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_23]), c_0_24]), c_0_29])])).
cnf(c_0_61, plain, (X1=k1_yellow_0(X2,X3)|~r1_orders_2(X2,X1,esk4_3(X2,X1,X3))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_62, negated_conjecture, (X1=k1_yellow_0(X2,k5_waybel_0(esk1_0,X3))|r1_orders_2(X2,X3,esk4_3(X2,X1,k5_waybel_0(esk1_0,X3)))|v2_struct_0(esk1_0)|~r2_lattice3(X2,k5_waybel_0(esk1_0,X3),X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_24])])).
cnf(c_0_63, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r2_hidden(X2,k5_waybel_0(X1,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_64, plain, (r2_hidden(esk3_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_65, negated_conjecture, (v2_struct_0(esk1_0)|v4_waybel_7(k1_yellow_0(esk1_0,k5_waybel_0(esk1_0,esk2_0)),esk1_0)|~m1_subset_1(k5_waybel_0(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_23]), c_0_24]), c_0_29])])).
cnf(c_0_66, negated_conjecture, (k1_yellow_0(X1,k5_waybel_0(esk1_0,X2))=X2|v2_struct_0(esk1_0)|~r2_lattice3(X1,k5_waybel_0(esk1_0,X2),X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_67, negated_conjecture, (~v4_waybel_7(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_68, plain, (r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk3_3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_69, plain, (r1_orders_2(X1,esk3_3(X2,k5_waybel_0(X1,X3),X4),X3)|r2_lattice3(X2,k5_waybel_0(X1,X3),X4)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X1)|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_70, negated_conjecture, (v2_struct_0(esk1_0)|~r2_lattice3(esk1_0,k5_waybel_0(esk1_0,esk2_0),esk2_0)|~m1_subset_1(k5_waybel_0(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_23]), c_0_24]), c_0_27])]), c_0_67])).
cnf(c_0_71, plain, (r2_lattice3(X1,k5_waybel_0(X1,X2),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
fof(c_0_72, plain, ![X44]:(~l1_orders_2(X44)|(~v1_lattice3(X44)|~v2_struct_0(X44))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattice3])])])).
cnf(c_0_73, negated_conjecture, (v2_struct_0(esk1_0)|~m1_subset_1(k5_waybel_0(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_23]), c_0_24])])).
cnf(c_0_74, plain, (~l1_orders_2(X1)|~v1_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_75, negated_conjecture, (v2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_52]), c_0_23]), c_0_24])])).
cnf(c_0_76, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_24]), c_0_26])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 77
# Proof object clause steps            : 48
# Proof object formula steps           : 29
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 52
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 41
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 41
# Processed clauses                    : 544
# ...of these trivial                  : 0
# ...subsumed                          : 148
# ...remaining for further processing  : 396
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 122
# Backward-rewritten                   : 92
# Generated clauses                    : 1578
# ...of the previous two non-trivial   : 1519
# Contextual simplify-reflections      : 135
# Paramodulations                      : 1560
# Factorizations                       : 0
# Equation resolutions                 : 18
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 141
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 130
# Current number of unprocessed clauses: 948
# ...number of literals in the above   : 11837
# Current number of archived formulas  : 0
# Current number of archived clauses   : 255
# Clause-clause subsumption calls (NU) : 37138
# Rec. Clause-clause subsumption calls : 1889
# Non-unit clause-clause subsumptions  : 341
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 124236

# -------------------------------------------------
# User time                : 0.081 s
# System time              : 0.000 s
# Total time               : 0.081 s
# Maximum resident set size: 28608 pages
