# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/1076fd6455bc4e1eae4000b2158622f0.E_input', t1_boole)).
fof(d2_xboole_0, axiom, k1_xboole_0=o_0_0_xboole_0, file('data/MPTP/proofs/1076fd6455bc4e1eae4000b2158622f0.E_input', d2_xboole_0)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/1076fd6455bc4e1eae4000b2158622f0.E_input', commutativity_k2_xboole_0)).
fof(t22_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/1076fd6455bc4e1eae4000b2158622f0.E_input', t22_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/1076fd6455bc4e1eae4000b2158622f0.E_input', t28_xboole_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/1076fd6455bc4e1eae4000b2158622f0.E_input', t36_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/1076fd6455bc4e1eae4000b2158622f0.E_input', commutativity_k3_xboole_0)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/1076fd6455bc4e1eae4000b2158622f0.E_input', d5_xboole_0)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/1076fd6455bc4e1eae4000b2158622f0.E_input', t4_xboole_0)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/1076fd6455bc4e1eae4000b2158622f0.E_input', idempotence_k3_xboole_0)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/1076fd6455bc4e1eae4000b2158622f0.E_input', t3_xboole_0)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/1076fd6455bc4e1eae4000b2158622f0.E_input', l32_xboole_1)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/1076fd6455bc4e1eae4000b2158622f0.E_input', t41_xboole_1)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/1076fd6455bc4e1eae4000b2158622f0.E_input', t39_xboole_1)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/1076fd6455bc4e1eae4000b2158622f0.E_input', t17_xboole_1)).
fof(t16_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3)), file('data/MPTP/proofs/1076fd6455bc4e1eae4000b2158622f0.E_input', t16_xboole_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/1076fd6455bc4e1eae4000b2158622f0.E_input', t7_xboole_0)).
fof(t32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k4_xboole_0(X2,X1)=>X1=X2), file('data/MPTP/proofs/1076fd6455bc4e1eae4000b2158622f0.E_input', t32_xboole_1)).
fof(t47_xboole_1, conjecture, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/1076fd6455bc4e1eae4000b2158622f0.E_input', t47_xboole_1)).
fof(c_0_19, plain, ![X25]:k2_xboole_0(X25,k1_xboole_0)=X25, inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_20, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_21, plain, (k1_xboole_0=o_0_0_xboole_0), inference(split_conjunct,[status(thm)],[d2_xboole_0])).
fof(c_0_22, plain, ![X9, X10]:k2_xboole_0(X9,X10)=k2_xboole_0(X10,X9), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_23, plain, (k2_xboole_0(X1,o_0_0_xboole_0)=X1), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_24, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_25, plain, ![X38, X39]:k2_xboole_0(X38,k3_xboole_0(X38,X39))=X38, inference(variable_rename,[status(thm)],[t22_xboole_1])).
fof(c_0_26, plain, ![X54, X55]:(~r1_tarski(X54,X55)|k3_xboole_0(X54,X55)=X54), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_27, plain, ![X14, X15]:r1_tarski(k4_xboole_0(X14,X15),X14), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_28, plain, ![X42, X43]:k3_xboole_0(X42,X43)=k3_xboole_0(X43,X42), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_29, plain, ![X26, X27, X28, X29, X30, X31, X32, X33]:((((r2_hidden(X29,X26)|~r2_hidden(X29,X28)|X28!=k4_xboole_0(X26,X27))&(~r2_hidden(X29,X27)|~r2_hidden(X29,X28)|X28!=k4_xboole_0(X26,X27)))&(~r2_hidden(X30,X26)|r2_hidden(X30,X27)|r2_hidden(X30,X28)|X28!=k4_xboole_0(X26,X27)))&((~r2_hidden(esk4_3(X31,X32,X33),X33)|(~r2_hidden(esk4_3(X31,X32,X33),X31)|r2_hidden(esk4_3(X31,X32,X33),X32))|X33=k4_xboole_0(X31,X32))&((r2_hidden(esk4_3(X31,X32,X33),X31)|r2_hidden(esk4_3(X31,X32,X33),X33)|X33=k4_xboole_0(X31,X32))&(~r2_hidden(esk4_3(X31,X32,X33),X32)|r2_hidden(esk4_3(X31,X32,X33),X33)|X33=k4_xboole_0(X31,X32))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_30, plain, (k2_xboole_0(o_0_0_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_31, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_34, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, plain, (k3_xboole_0(o_0_0_xboole_0,X1)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_37, plain, (k3_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])).
fof(c_0_38, plain, ![X44, X45, X47, X48, X49]:((r1_xboole_0(X44,X45)|r2_hidden(esk5_2(X44,X45),k3_xboole_0(X44,X45)))&(~r2_hidden(X49,k3_xboole_0(X47,X48))|~r1_xboole_0(X47,X48))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_39, plain, ![X37]:k3_xboole_0(X37,X37)=X37, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_40, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_35])).
cnf(c_0_41, plain, (k4_xboole_0(o_0_0_xboole_0,X1)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_42, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk5_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_43, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_44, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_45, plain, ![X16, X17, X19, X20, X21]:(((r2_hidden(esk3_2(X16,X17),X16)|r1_xboole_0(X16,X17))&(r2_hidden(esk3_2(X16,X17),X17)|r1_xboole_0(X16,X17)))&(~r2_hidden(X21,X19)|~r2_hidden(X21,X20)|~r1_xboole_0(X19,X20))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_46, plain, (~r2_hidden(X1,o_0_0_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_47, plain, (r2_hidden(esk5_2(X1,X1),X1)|r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_48, plain, (k3_xboole_0(X1,o_0_0_xboole_0)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_34, c_0_36])).
cnf(c_0_49, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_44, c_0_43])).
cnf(c_0_50, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_51, plain, (r1_xboole_0(o_0_0_xboole_0,o_0_0_xboole_0)|~r2_hidden(esk5_2(o_0_0_xboole_0,o_0_0_xboole_0),X1)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_52, plain, (r2_hidden(esk5_2(X1,o_0_0_xboole_0),o_0_0_xboole_0)|r1_xboole_0(X1,o_0_0_xboole_0)), inference(spm,[status(thm)],[c_0_42, c_0_48])).
cnf(c_0_53, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_54, plain, (r1_xboole_0(o_0_0_xboole_0,o_0_0_xboole_0)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_55, plain, (~r2_hidden(X1,o_0_0_xboole_0)|~r1_xboole_0(o_0_0_xboole_0,X2)), inference(spm,[status(thm)],[c_0_44, c_0_36])).
cnf(c_0_56, plain, (r1_xboole_0(o_0_0_xboole_0,X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_57, plain, (r2_hidden(esk4_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk4_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_58, plain, (r2_hidden(esk4_3(X1,X2,X3),X1)|r2_hidden(esk4_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_59, plain, ![X7, X8]:((k4_xboole_0(X7,X8)!=k1_xboole_0|r1_tarski(X7,X8))&(~r1_tarski(X7,X8)|k4_xboole_0(X7,X8)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
fof(c_0_60, plain, ![X11, X12, X13]:k4_xboole_0(k4_xboole_0(X11,X12),X13)=k4_xboole_0(X11,k2_xboole_0(X12,X13)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
cnf(c_0_61, plain, (~r2_hidden(X1,o_0_0_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_56])])).
cnf(c_0_62, plain, (X1=k4_xboole_0(X2,X2)|r2_hidden(esk4_3(X2,X2,X1),X1)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
fof(c_0_63, plain, ![X40, X41]:k2_xboole_0(X40,k4_xboole_0(X41,X40))=k2_xboole_0(X40,X41), inference(variable_rename,[status(thm)],[t39_xboole_1])).
fof(c_0_64, plain, ![X50, X51]:r1_tarski(k3_xboole_0(X50,X51),X50), inference(variable_rename,[status(thm)],[t17_xboole_1])).
cnf(c_0_65, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_66, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_67, plain, (k4_xboole_0(X1,X1)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_68, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_69, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
fof(c_0_70, plain, ![X22, X23, X24]:k3_xboole_0(k3_xboole_0(X22,X23),X24)=k3_xboole_0(X22,k3_xboole_0(X23,X24)), inference(variable_rename,[status(thm)],[t16_xboole_1])).
cnf(c_0_71, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=o_0_0_xboole_0), inference(rw,[status(thm)],[c_0_65, c_0_21])).
cnf(c_0_72, plain, (k4_xboole_0(X1,k2_xboole_0(X2,X1))=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68])).
cnf(c_0_73, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(X2,X3)),k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_33, c_0_66])).
cnf(c_0_74, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_69, c_0_34])).
cnf(c_0_75, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_76, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_77, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(X2,X3)),k4_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_73, c_0_24])).
cnf(c_0_78, plain, (r1_tarski(k3_xboole_0(X1,k3_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_79, plain, (k3_xboole_0(X1,k2_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_32, c_0_76])).
cnf(c_0_80, plain, (r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_77, c_0_31])).
cnf(c_0_81, plain, (r1_tarski(k3_xboole_0(X1,X2),k2_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_82, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X2,X3)))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_32, c_0_80])).
cnf(c_0_83, plain, (r1_tarski(X1,k2_xboole_0(X2,k2_xboole_0(X3,X1)))), inference(spm,[status(thm)],[c_0_81, c_0_79])).
cnf(c_0_84, plain, (k2_xboole_0(X1,k4_xboole_0(X1,X2))=X1), inference(spm,[status(thm)],[c_0_31, c_0_37])).
cnf(c_0_85, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|k4_xboole_0(X1,k2_xboole_0(X2,X3))!=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_71, c_0_66])).
cnf(c_0_86, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X1)=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_67]), c_0_48])).
cnf(c_0_87, plain, (r1_tarski(k4_xboole_0(X1,X2),k2_xboole_0(X3,X1))), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_88, plain, (k2_xboole_0(X1,k3_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_31, c_0_34])).
cnf(c_0_89, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(k2_xboole_0(X1,X2),X3),X1),X2)), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_90, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
fof(c_0_91, plain, ![X52]:(X52=k1_xboole_0|r2_hidden(esk6_1(X52),X52)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_92, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_93, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(k2_xboole_0(X1,X2),X3),X2),X1)), inference(spm,[status(thm)],[c_0_89, c_0_24])).
cnf(c_0_94, plain, (k3_xboole_0(X1,k4_xboole_0(k3_xboole_0(X2,X1),X3))=k4_xboole_0(k3_xboole_0(X2,X1),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_90]), c_0_34])).
cnf(c_0_95, plain, (X1=k1_xboole_0|r2_hidden(esk6_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_96, plain, (r2_hidden(esk3_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_97, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_92])).
cnf(c_0_98, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(k2_xboole_0(X1,X2),X3),k4_xboole_0(X2,X1)),X1)), inference(spm,[status(thm)],[c_0_93, c_0_68])).
cnf(c_0_99, plain, (k3_xboole_0(k2_xboole_0(X1,X2),k4_xboole_0(X2,X3))=k4_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_94, c_0_79])).
cnf(c_0_100, plain, (r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X3,X2)))), inference(spm,[status(thm)],[c_0_77, c_0_88])).
cnf(c_0_101, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_44, c_0_34])).
cnf(c_0_102, plain, (X1=o_0_0_xboole_0|r2_hidden(esk6_1(X1),X1)), inference(rw,[status(thm)],[c_0_95, c_0_21])).
cnf(c_0_103, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r2_hidden(esk3_2(X1,k4_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_40, c_0_96])).
cnf(c_0_104, plain, (r2_hidden(esk3_2(k4_xboole_0(X1,X2),X3),X1)|r1_xboole_0(k4_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_97, c_0_50])).
cnf(c_0_105, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(X2,k4_xboole_0(X1,X3))),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_66])).
cnf(c_0_106, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X3,X2)))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_32, c_0_100])).
fof(c_0_107, plain, ![X35, X36]:(k4_xboole_0(X35,X36)!=k4_xboole_0(X36,X35)|X35=X36), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t32_xboole_1])])).
cnf(c_0_108, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_109, plain, (k3_xboole_0(X1,X2)=o_0_0_xboole_0|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_110, plain, (r1_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X3,X1))), inference(spm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_111, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_105, c_0_30])).
cnf(c_0_112, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X2)=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_67]), c_0_48])).
cnf(c_0_113, plain, (X1=X2|k4_xboole_0(X1,X2)!=k4_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_114, plain, (k4_xboole_0(X1,X2)=o_0_0_xboole_0|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_108, c_0_21])).
cnf(c_0_115, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X3))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_109, c_0_110])).
cnf(c_0_116, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X1)))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_111]), c_0_34])).
cnf(c_0_117, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,k2_xboole_0(X2,X3)),X2),X3)), inference(spm,[status(thm)],[c_0_85, c_0_112])).
cnf(c_0_118, plain, (X1=X2|k4_xboole_0(X1,X2)!=o_0_0_xboole_0|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_113, c_0_114])).
cnf(c_0_119, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,X1)))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_115, c_0_116])).
cnf(c_0_120, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,k2_xboole_0(X2,X3)),X2),k4_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_117, c_0_68])).
cnf(c_0_121, plain, (k3_xboole_0(X1,k2_xboole_0(X2,k2_xboole_0(X3,X1)))=X1), inference(spm,[status(thm)],[c_0_32, c_0_83])).
cnf(c_0_122, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_33])])).
cnf(c_0_123, plain, (r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(k2_xboole_0(X3,X1),X2))), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_124, plain, (k4_xboole_0(X1,k3_xboole_0(X2,k4_xboole_0(X3,X1)))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_122]), c_0_37])).
cnf(c_0_125, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,X2),X3),k4_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_123, c_0_88])).
cnf(c_0_126, plain, (k4_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X3,X1))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_124, c_0_82])).
cnf(c_0_127, plain, (r1_tarski(k3_xboole_0(X1,X2),k4_xboole_0(X2,k4_xboole_0(X3,X1)))), inference(spm,[status(thm)],[c_0_125, c_0_126])).
cnf(c_0_128, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X3,X1)))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_127]), c_0_75]), c_0_37])).
fof(c_0_129, negated_conjecture, ~(![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(assume_negation,[status(cth)],[t47_xboole_1])).
cnf(c_0_130, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X2,X1)), inference(rw,[status(thm)],[c_0_116, c_0_128])).
fof(c_0_131, negated_conjecture, k4_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0))!=k4_xboole_0(esk1_0,esk2_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_129])])])).
cnf(c_0_132, plain, (k4_xboole_0(X1,k3_xboole_0(X2,X1))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_130]), c_0_34]), c_0_37])).
cnf(c_0_133, negated_conjecture, (k4_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0))!=k4_xboole_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_131])).
cnf(c_0_134, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_132, c_0_34])).
cnf(c_0_135, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_133, c_0_134])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 136
# Proof object clause steps            : 98
# Proof object formula steps           : 38
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 64
# Proof object simplifying inferences  : 25
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 28
# Processed clauses                    : 8182
# ...of these trivial                  : 2181
# ...subsumed                          : 4743
# ...remaining for further processing  : 1258
# Other redundant clauses eliminated   : 708
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 11
# Backward-rewritten                   : 87
# Generated clauses                    : 254610
# ...of the previous two non-trivial   : 115180
# Contextual simplify-reflections      : 2
# Paramodulations                      : 253895
# Factorizations                       : 6
# Equation resolutions                 : 709
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1157
#    Positive orientable unit clauses  : 852
#    Positive unorientable unit clauses: 5
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 297
# Current number of unprocessed clauses: 106612
# ...number of literals in the above   : 169697
# Current number of archived formulas  : 0
# Current number of archived clauses   : 98
# Clause-clause subsumption calls (NU) : 38843
# Rec. Clause-clause subsumption calls : 34891
# Non-unit clause-clause subsumptions  : 2512
# Unit Clause-clause subsumption calls : 889
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7811
# BW rewrite match successes           : 169
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2020075

# -------------------------------------------------
# User time                : 1.248 s
# System time              : 0.027 s
# Total time               : 1.275 s
# Maximum resident set size: 91884 pages
