# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/109b86602f6e450ab16fc51f05d24b20.E_input', t36_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/109b86602f6e450ab16fc51f05d24b20.E_input', t3_boole)).
fof(t111_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k3_xboole_0(k4_xboole_0(X1,X3),X2), file('data/MPTP/proofs/109b86602f6e450ab16fc51f05d24b20.E_input', t111_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/109b86602f6e450ab16fc51f05d24b20.E_input', t48_xboole_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/109b86602f6e450ab16fc51f05d24b20.E_input', l32_xboole_1)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/109b86602f6e450ab16fc51f05d24b20.E_input', d2_tarski)).
fof(t77_enumset1, axiom, ![X1, X2]:k2_enumset1(X1,X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/109b86602f6e450ab16fc51f05d24b20.E_input', t77_enumset1)).
fof(t90_enumset1, axiom, ![X1, X2, X3, X4]:k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/109b86602f6e450ab16fc51f05d24b20.E_input', t90_enumset1)).
fof(t42_zfmisc_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_tarski(X2,X3))<=>~((~(X1=k1_xboole_0)&(~(X1=k1_tarski(X2))&(~(X1=k1_tarski(X3))&~(X1=k2_tarski(X2,X3))))))), file('data/MPTP/proofs/109b86602f6e450ab16fc51f05d24b20.E_input', t42_zfmisc_1)).
fof(t50_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3)), file('data/MPTP/proofs/109b86602f6e450ab16fc51f05d24b20.E_input', t50_xboole_1)).
fof(t123_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X2,X3,X1), file('data/MPTP/proofs/109b86602f6e450ab16fc51f05d24b20.E_input', t123_enumset1)).
fof(t47_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/109b86602f6e450ab16fc51f05d24b20.E_input', t47_xboole_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/109b86602f6e450ab16fc51f05d24b20.E_input', d3_tarski)).
fof(t52_zfmisc_1, conjecture, ![X1, X2]:(r2_hidden(X1,X2)=>k3_xboole_0(X2,k1_tarski(X1))=k1_tarski(X1)), file('data/MPTP/proofs/109b86602f6e450ab16fc51f05d24b20.E_input', t52_zfmisc_1)).
fof(t32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k4_xboole_0(X2,X1)=>X1=X2), file('data/MPTP/proofs/109b86602f6e450ab16fc51f05d24b20.E_input', t32_xboole_1)).
fof(c_0_15, plain, ![X5, X6]:r1_tarski(k4_xboole_0(X5,X6),X5), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_16, plain, ![X7]:k4_xboole_0(X7,k1_xboole_0)=X7, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_17, plain, ![X8, X9, X10]:k4_xboole_0(k3_xboole_0(X8,X9),k3_xboole_0(X10,X9))=k3_xboole_0(k4_xboole_0(X8,X10),X9), inference(variable_rename,[status(thm)],[t111_xboole_1])).
fof(c_0_18, plain, ![X11, X12]:k4_xboole_0(X11,k4_xboole_0(X11,X12))=k3_xboole_0(X11,X12), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_19, plain, ![X13, X14]:((k4_xboole_0(X13,X14)!=k1_xboole_0|r1_tarski(X13,X14))&(~r1_tarski(X13,X14)|k4_xboole_0(X13,X14)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_20, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, plain, ![X15, X16, X17, X18, X19, X20, X21, X22]:(((~r2_hidden(X18,X17)|(X18=X15|X18=X16)|X17!=k2_tarski(X15,X16))&((X19!=X15|r2_hidden(X19,X17)|X17!=k2_tarski(X15,X16))&(X19!=X16|r2_hidden(X19,X17)|X17!=k2_tarski(X15,X16))))&(((esk1_3(X20,X21,X22)!=X20|~r2_hidden(esk1_3(X20,X21,X22),X22)|X22=k2_tarski(X20,X21))&(esk1_3(X20,X21,X22)!=X21|~r2_hidden(esk1_3(X20,X21,X22),X22)|X22=k2_tarski(X20,X21)))&(r2_hidden(esk1_3(X20,X21,X22),X22)|(esk1_3(X20,X21,X22)=X20|esk1_3(X20,X21,X22)=X21)|X22=k2_tarski(X20,X21)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_23, plain, ![X24, X25]:k2_enumset1(X24,X24,X24,X25)=k2_tarski(X24,X25), inference(variable_rename,[status(thm)],[t77_enumset1])).
fof(c_0_24, plain, ![X26, X27, X28, X29]:k6_enumset1(X26,X26,X26,X26,X26,X27,X28,X29)=k2_enumset1(X26,X27,X28,X29), inference(variable_rename,[status(thm)],[t90_enumset1])).
fof(c_0_25, plain, ![X30, X31, X32]:((~r1_tarski(X30,k2_tarski(X31,X32))|(X30=k1_xboole_0|(X30=k1_tarski(X31)|(X30=k1_tarski(X32)|X30=k2_tarski(X31,X32)))))&((X30!=k1_xboole_0|r1_tarski(X30,k2_tarski(X31,X32)))&((X30!=k1_tarski(X31)|r1_tarski(X30,k2_tarski(X31,X32)))&((X30!=k1_tarski(X32)|r1_tarski(X30,k2_tarski(X31,X32)))&(X30!=k2_tarski(X31,X32)|r1_tarski(X30,k2_tarski(X31,X32))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t42_zfmisc_1])])])])).
fof(c_0_26, plain, ![X33, X34, X35]:k3_xboole_0(X33,k4_xboole_0(X34,X35))=k4_xboole_0(k3_xboole_0(X33,X34),k3_xboole_0(X33,X35)), inference(variable_rename,[status(thm)],[t50_xboole_1])).
cnf(c_0_27, plain, (k4_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k3_xboole_0(k4_xboole_0(X1,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_29, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_30, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_31, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32, plain, (k2_enumset1(X1,X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_33, plain, (k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_34, plain, ![X36, X37, X38, X39]:k2_enumset1(X36,X37,X38,X39)=k2_enumset1(X39,X37,X38,X36), inference(variable_rename,[status(thm)],[t123_enumset1])).
cnf(c_0_35, plain, (r1_tarski(X1,k2_tarski(X3,X2))|X1!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_36, plain, ![X40, X41]:k4_xboole_0(X40,k3_xboole_0(X40,X41))=k4_xboole_0(X40,X41), inference(variable_rename,[status(thm)],[t47_xboole_1])).
cnf(c_0_37, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_38, plain, (k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k4_xboole_0(X3,k4_xboole_0(X3,X2)))=k4_xboole_0(k4_xboole_0(X1,X3),k4_xboole_0(k4_xboole_0(X1,X3),X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_28]), c_0_28]), c_0_28])).
cnf(c_0_39, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_40, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_29, c_0_20])).
cnf(c_0_41, plain, (X1=X4|X1=X3|X2!=k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_32]), c_0_33])).
cnf(c_0_42, plain, (k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_43, plain, ![X42, X43, X44, X45, X46]:((~r1_tarski(X42,X43)|(~r2_hidden(X44,X42)|r2_hidden(X44,X43)))&((r2_hidden(esk2_2(X45,X46),X45)|r1_tarski(X45,X46))&(~r2_hidden(esk2_2(X45,X46),X46)|r1_tarski(X45,X46)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_44, plain, (r1_tarski(X1,k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X2))|X1!=k1_tarski(X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_32]), c_0_33])).
cnf(c_0_45, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_46, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,X3)))=k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k4_xboole_0(X1,k4_xboole_0(X1,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_28]), c_0_28]), c_0_28])).
cnf(c_0_47, plain, (k4_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X1)))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_21]), c_0_40]), c_0_21])).
cnf(c_0_48, plain, (X1=X2|X1=X3|~r2_hidden(X1,k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X2))), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_49, plain, (k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k6_enumset1(X4,X4,X4,X4,X4,X2,X3,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_33]), c_0_33])).
cnf(c_0_50, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_51, plain, (r1_tarski(k1_tarski(X1),k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X1))), inference(er,[status(thm)],[c_0_44])).
fof(c_0_52, negated_conjecture, ~(![X1, X2]:(r2_hidden(X1,X2)=>k3_xboole_0(X2,k1_tarski(X1))=k1_tarski(X1))), inference(assume_negation,[status(cth)],[t52_zfmisc_1])).
fof(c_0_53, plain, ![X50, X51]:(k4_xboole_0(X50,X51)!=k4_xboole_0(X51,X50)|X50=X51), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t32_xboole_1])])).
cnf(c_0_54, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X1,X2)))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_45, c_0_28])).
cnf(c_0_55, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,X1)))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_40]), c_0_47])).
cnf(c_0_56, plain, (X1=X2|X1=X3|~r2_hidden(X1,k6_enumset1(X2,X2,X2,X2,X2,X3,X3,X3))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_57, plain, (r2_hidden(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3))|~r2_hidden(X1,k1_tarski(X3))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
fof(c_0_58, negated_conjecture, (r2_hidden(esk3_0,esk4_0)&k3_xboole_0(esk4_0,k1_tarski(esk3_0))!=k1_tarski(esk3_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_52])])])).
cnf(c_0_59, plain, (X1=X2|k4_xboole_0(X1,X2)!=k4_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_60, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_21])).
cnf(c_0_61, plain, (X1=X2|~r2_hidden(X1,k1_tarski(X2))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_62, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_63, negated_conjecture, (k3_xboole_0(esk4_0,k1_tarski(esk3_0))!=k1_tarski(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_64, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X3,k4_xboole_0(X3,X2))|k4_xboole_0(k4_xboole_0(X1,X3),k4_xboole_0(k4_xboole_0(X1,X3),X2))!=k4_xboole_0(k4_xboole_0(X3,X1),k4_xboole_0(k4_xboole_0(X3,X1),X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_38]), c_0_38])).
cnf(c_0_65, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_60]), c_0_39]), c_0_21])).
cnf(c_0_66, plain, (r1_tarski(X1,X2)|~r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_67, plain, (esk2_2(k1_tarski(X1),X2)=X1|r1_tarski(k1_tarski(X1),X2)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_68, negated_conjecture, (k4_xboole_0(esk4_0,k4_xboole_0(esk4_0,k1_tarski(esk3_0)))!=k1_tarski(esk3_0)), inference(rw,[status(thm)],[c_0_63, c_0_28])).
cnf(c_0_69, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=X2|k4_xboole_0(X2,X1)!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_40]), c_0_39]), c_0_21]), c_0_65]), c_0_39]), c_0_21])).
cnf(c_0_70, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_71, negated_conjecture, (k4_xboole_0(k1_tarski(esk3_0),esk4_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_72, plain, (k4_xboole_0(k1_tarski(X1),X2)=k1_xboole_0|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_29, c_0_70])).
cnf(c_0_73, negated_conjecture, (r2_hidden(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_74, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 75
# Proof object clause steps            : 44
# Proof object formula steps           : 31
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 32
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 25
# Processed clauses                    : 244
# ...of these trivial                  : 11
# ...subsumed                          : 120
# ...remaining for further processing  : 113
# Other redundant clauses eliminated   : 68
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 6
# Generated clauses                    : 2738
# ...of the previous two non-trivial   : 1808
# Contextual simplify-reflections      : 0
# Paramodulations                      : 2667
# Factorizations                       : 2
# Equation resolutions                 : 71
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 75
#    Positive orientable unit clauses  : 29
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 41
# Current number of unprocessed clauses: 1596
# ...number of literals in the above   : 3054
# Current number of archived formulas  : 0
# Current number of archived clauses   : 34
# Clause-clause subsumption calls (NU) : 383
# Rec. Clause-clause subsumption calls : 381
# Non-unit clause-clause subsumptions  : 119
# Unit Clause-clause subsumption calls : 9
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 133
# BW rewrite match successes           : 18
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 51201

# -------------------------------------------------
# User time                : 0.032 s
# System time              : 0.000 s
# Total time               : 0.032 s
# Maximum resident set size: 28648 pages
