# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/10be12ecdb3c4b09bdf9fa8cd36889fc.E_input', t12_setfam_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/10be12ecdb3c4b09bdf9fa8cd36889fc.E_input', t70_enumset1)).
fof(d5_tsep_1, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(~(r1_tsep_1(X2,X3))=>![X4]:((~(v2_struct_0(X4))&(v1_pre_topc(X4)&m1_pre_topc(X4,X1)))=>(X4=k2_tsep_1(X1,X2,X3)<=>u1_struct_0(X4)=k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)))))))), file('data/MPTP/proofs/10be12ecdb3c4b09bdf9fa8cd36889fc.E_input', d5_tsep_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/10be12ecdb3c4b09bdf9fa8cd36889fc.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/10be12ecdb3c4b09bdf9fa8cd36889fc.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/10be12ecdb3c4b09bdf9fa8cd36889fc.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/10be12ecdb3c4b09bdf9fa8cd36889fc.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/10be12ecdb3c4b09bdf9fa8cd36889fc.E_input', t75_enumset1)).
fof(dt_k2_tsep_1, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))&(~(v2_struct_0(X3))&m1_pre_topc(X3,X1))))=>(~(v2_struct_0(k2_tsep_1(X1,X2,X3)))&(v1_pre_topc(k2_tsep_1(X1,X2,X3))&m1_pre_topc(k2_tsep_1(X1,X2,X3),X1)))), file('data/MPTP/proofs/10be12ecdb3c4b09bdf9fa8cd36889fc.E_input', dt_k2_tsep_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/10be12ecdb3c4b09bdf9fa8cd36889fc.E_input', t3_subset)).
fof(dt_k9_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>m1_subset_1(k9_subset_1(X1,X2,X3),k1_zfmisc_1(X1))), file('data/MPTP/proofs/10be12ecdb3c4b09bdf9fa8cd36889fc.E_input', dt_k9_subset_1)).
fof(redefinition_k9_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>k9_subset_1(X1,X2,X3)=k3_xboole_0(X2,X3)), file('data/MPTP/proofs/10be12ecdb3c4b09bdf9fa8cd36889fc.E_input', redefinition_k9_subset_1)).
fof(t17_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(~(r1_tsep_1(X2,X3))=>![X4]:(m1_subset_1(X4,u1_struct_0(k2_tsep_1(X1,X2,X3)))=>(?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&X5=X4)&?[X5]:(m1_subset_1(X5,u1_struct_0(X3))&X5=X4))))))), file('data/MPTP/proofs/10be12ecdb3c4b09bdf9fa8cd36889fc.E_input', t17_tmap_1)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/10be12ecdb3c4b09bdf9fa8cd36889fc.E_input', t17_xboole_1)).
fof(t4_tsep_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X1)=>(r1_tarski(u1_struct_0(X2),u1_struct_0(X3))<=>m1_pre_topc(X2,X3))))), file('data/MPTP/proofs/10be12ecdb3c4b09bdf9fa8cd36889fc.E_input', t4_tsep_1)).
fof(t55_pre_topc, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>m1_subset_1(X3,u1_struct_0(X1))))), file('data/MPTP/proofs/10be12ecdb3c4b09bdf9fa8cd36889fc.E_input', t55_pre_topc)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/10be12ecdb3c4b09bdf9fa8cd36889fc.E_input', dt_m1_pre_topc)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/10be12ecdb3c4b09bdf9fa8cd36889fc.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/10be12ecdb3c4b09bdf9fa8cd36889fc.E_input', d4_subset_1)).
fof(c_0_19, plain, ![X8, X9]:k1_setfam_1(k2_tarski(X8,X9))=k3_xboole_0(X8,X9), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_20, plain, ![X10, X11]:k1_enumset1(X10,X10,X11)=k2_tarski(X10,X11), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_21, plain, ![X12, X13, X14, X15]:((X15!=k2_tsep_1(X12,X13,X14)|u1_struct_0(X15)=k3_xboole_0(u1_struct_0(X13),u1_struct_0(X14))|(v2_struct_0(X15)|(~v1_pre_topc(X15)|~m1_pre_topc(X15,X12)))|r1_tsep_1(X13,X14)|(v2_struct_0(X14)|~m1_pre_topc(X14,X12))|(v2_struct_0(X13)|~m1_pre_topc(X13,X12))|(v2_struct_0(X12)|~l1_pre_topc(X12)))&(u1_struct_0(X15)!=k3_xboole_0(u1_struct_0(X13),u1_struct_0(X14))|X15=k2_tsep_1(X12,X13,X14)|(v2_struct_0(X15)|(~v1_pre_topc(X15)|~m1_pre_topc(X15,X12)))|r1_tsep_1(X13,X14)|(v2_struct_0(X14)|~m1_pre_topc(X14,X12))|(v2_struct_0(X13)|~m1_pre_topc(X13,X12))|(v2_struct_0(X12)|~l1_pre_topc(X12)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_tsep_1])])])])])).
cnf(c_0_22, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_24, plain, ![X16, X17, X18]:k2_enumset1(X16,X16,X17,X18)=k1_enumset1(X16,X17,X18), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_25, plain, ![X19, X20, X21, X22]:k3_enumset1(X19,X19,X20,X21,X22)=k2_enumset1(X19,X20,X21,X22), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_26, plain, ![X23, X24, X25, X26, X27]:k4_enumset1(X23,X23,X24,X25,X26,X27)=k3_enumset1(X23,X24,X25,X26,X27), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_27, plain, ![X28, X29, X30, X31, X32, X33]:k5_enumset1(X28,X28,X29,X30,X31,X32,X33)=k4_enumset1(X28,X29,X30,X31,X32,X33), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_28, plain, ![X34, X35, X36, X37, X38, X39, X40]:k6_enumset1(X34,X34,X35,X36,X37,X38,X39,X40)=k5_enumset1(X34,X35,X36,X37,X38,X39,X40), inference(variable_rename,[status(thm)],[t75_enumset1])).
cnf(c_0_29, plain, (u1_struct_0(X1)=k3_xboole_0(u1_struct_0(X3),u1_struct_0(X4))|v2_struct_0(X1)|r1_tsep_1(X3,X4)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|X1!=k2_tsep_1(X2,X3,X4)|~v1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_31, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_36, plain, ![X41, X42, X43]:((~v2_struct_0(k2_tsep_1(X41,X42,X43))|(v2_struct_0(X41)|~l1_pre_topc(X41)|(v2_struct_0(X42)|~m1_pre_topc(X42,X41)|(v2_struct_0(X43)|~m1_pre_topc(X43,X41)))))&((v1_pre_topc(k2_tsep_1(X41,X42,X43))|(v2_struct_0(X41)|~l1_pre_topc(X41)|(v2_struct_0(X42)|~m1_pre_topc(X42,X41)|(v2_struct_0(X43)|~m1_pre_topc(X43,X41)))))&(m1_pre_topc(k2_tsep_1(X41,X42,X43),X41)|(v2_struct_0(X41)|~l1_pre_topc(X41)|(v2_struct_0(X42)|~m1_pre_topc(X42,X41)|(v2_struct_0(X43)|~m1_pre_topc(X43,X41))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_tsep_1])])])])).
fof(c_0_37, plain, ![X44, X45]:((~m1_subset_1(X44,k1_zfmisc_1(X45))|r1_tarski(X44,X45))&(~r1_tarski(X44,X45)|m1_subset_1(X44,k1_zfmisc_1(X45)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_38, plain, ![X46, X47, X48]:(~m1_subset_1(X48,k1_zfmisc_1(X46))|m1_subset_1(k9_subset_1(X46,X47,X48),k1_zfmisc_1(X46))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k9_subset_1])])).
fof(c_0_39, plain, ![X49, X50, X51]:(~m1_subset_1(X51,k1_zfmisc_1(X49))|k9_subset_1(X49,X50,X51)=k3_xboole_0(X50,X51)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k9_subset_1])])).
cnf(c_0_40, plain, (u1_struct_0(X1)=k1_setfam_1(k6_enumset1(u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X4)))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|r1_tsep_1(X3,X4)|X1!=k2_tsep_1(X2,X3,X4)|~l1_pre_topc(X2)|~v1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35])).
cnf(c_0_41, plain, (m1_pre_topc(k2_tsep_1(X1,X2,X3),X1)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_42, plain, (v1_pre_topc(k2_tsep_1(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_43, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~v2_struct_0(k2_tsep_1(X1,X2,X3))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_44, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(~(r1_tsep_1(X2,X3))=>![X4]:(m1_subset_1(X4,u1_struct_0(k2_tsep_1(X1,X2,X3)))=>(?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&X5=X4)&?[X5]:(m1_subset_1(X5,u1_struct_0(X3))&X5=X4)))))))), inference(assume_negation,[status(cth)],[t17_tmap_1])).
cnf(c_0_45, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_46, plain, (m1_subset_1(k9_subset_1(X2,X3,X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_47, plain, (k9_subset_1(X2,X3,X1)=k3_xboole_0(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_48, plain, (u1_struct_0(k2_tsep_1(X1,X2,X3))=k1_setfam_1(k6_enumset1(u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X3)))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_40]), c_0_41]), c_0_42]), c_0_43])).
fof(c_0_49, negated_conjecture, ![X56, X57]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk1_0))&(~r1_tsep_1(esk2_0,esk3_0)&(m1_subset_1(esk4_0,u1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0)))&(~m1_subset_1(X56,u1_struct_0(esk2_0))|X56!=esk4_0|(~m1_subset_1(X57,u1_struct_0(esk3_0))|X57!=esk4_0))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_44])])])])])).
fof(c_0_50, plain, ![X58, X59]:r1_tarski(k3_xboole_0(X58,X59),X58), inference(variable_rename,[status(thm)],[t17_xboole_1])).
cnf(c_0_51, plain, (r1_tarski(k9_subset_1(X1,X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_52, plain, (k9_subset_1(X2,X3,X1)=k1_setfam_1(k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35])).
cnf(c_0_53, plain, (u1_struct_0(k2_tsep_1(X1,X2,X3))=u1_struct_0(k2_tsep_1(X4,X2,X3))|r1_tsep_1(X2,X3)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X4)|~m1_pre_topc(X2,X4)|~l1_pre_topc(X1)|~l1_pre_topc(X4)), inference(spm,[status(thm)],[c_0_48, c_0_48])).
cnf(c_0_54, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_55, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_56, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_57, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_58, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_59, plain, (r1_tarski(k1_setfam_1(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2)),X3)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_60, negated_conjecture, (u1_struct_0(k2_tsep_1(esk1_0,X1,esk3_0))=u1_struct_0(k2_tsep_1(X2,X1,esk3_0))|r1_tsep_1(X1,esk3_0)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(esk3_0,X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55])]), c_0_56]), c_0_57])).
cnf(c_0_61, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_62, negated_conjecture, (~r1_tsep_1(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_63, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_64, plain, ![X60, X61, X62]:((~r1_tarski(u1_struct_0(X61),u1_struct_0(X62))|m1_pre_topc(X61,X62)|~m1_pre_topc(X62,X60)|~m1_pre_topc(X61,X60)|(~v2_pre_topc(X60)|~l1_pre_topc(X60)))&(~m1_pre_topc(X61,X62)|r1_tarski(u1_struct_0(X61),u1_struct_0(X62))|~m1_pre_topc(X62,X60)|~m1_pre_topc(X61,X60)|(~v2_pre_topc(X60)|~l1_pre_topc(X60)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_tsep_1])])])])).
cnf(c_0_65, plain, (r1_tarski(k1_setfam_1(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2)),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35])).
cnf(c_0_66, plain, (r1_tarski(u1_struct_0(k2_tsep_1(X1,X2,X3)),X4)|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(u1_struct_0(X3),k1_zfmisc_1(X4))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_59, c_0_48])).
cnf(c_0_67, negated_conjecture, (u1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0))=u1_struct_0(k2_tsep_1(X1,esk2_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62]), c_0_63])).
fof(c_0_68, plain, ![X63, X64, X65]:(v2_struct_0(X63)|~l1_pre_topc(X63)|(v2_struct_0(X64)|~m1_pre_topc(X64,X63)|(~m1_subset_1(X65,u1_struct_0(X64))|m1_subset_1(X65,u1_struct_0(X63))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t55_pre_topc])])])])).
cnf(c_0_69, plain, (m1_pre_topc(X1,X2)|~r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_70, plain, (r1_tarski(u1_struct_0(k2_tsep_1(X1,X2,X3)),u1_struct_0(X2))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_65, c_0_48])).
fof(c_0_71, plain, ![X66, X67]:(~l1_pre_topc(X66)|(~m1_pre_topc(X67,X66)|l1_pre_topc(X67))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_72, negated_conjecture, (r1_tarski(u1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0)),X1)|v2_struct_0(X2)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(X1))|~m1_pre_topc(esk3_0,X2)|~m1_pre_topc(esk2_0,X2)|~l1_pre_topc(X2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_62]), c_0_56]), c_0_63])).
fof(c_0_73, plain, ![X68]:m1_subset_1(k2_subset_1(X68),k1_zfmisc_1(X68)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_74, plain, ![X69]:k2_subset_1(X69)=X69, inference(variable_rename,[status(thm)],[d4_subset_1])).
cnf(c_0_75, plain, (v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(X3,u1_struct_0(X1))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_76, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_77, plain, (r1_tsep_1(X1,X2)|m1_pre_topc(k2_tsep_1(X3,X1,X2),X1)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~v2_pre_topc(X4)|~m1_pre_topc(k2_tsep_1(X3,X1,X2),X4)|~m1_pre_topc(X1,X4)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X4)|~l1_pre_topc(X3)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_78, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_79, negated_conjecture, (r1_tarski(u1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0)),X1)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(X1))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_55]), c_0_54]), c_0_61])]), c_0_57])).
cnf(c_0_80, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_81, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_82, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(X1))|v2_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k2_tsep_1(esk1_0,esk2_0,esk3_0),X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_83, plain, (r1_tsep_1(X1,X2)|m1_pre_topc(k2_tsep_1(X3,X1,X2),X1)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|~v2_pre_topc(X3)|~m1_pre_topc(X1,X3)|~m1_pre_topc(X2,X3)|~l1_pre_topc(X3)), inference(spm,[status(thm)],[c_0_77, c_0_41])).
cnf(c_0_84, negated_conjecture, (l1_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_61]), c_0_55])])).
cnf(c_0_85, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_86, negated_conjecture, (m1_pre_topc(k2_tsep_1(esk1_0,esk2_0,esk3_0),X1)|~v2_pre_topc(X2)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_pre_topc(k2_tsep_1(esk1_0,esk2_0,esk3_0),X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_69, c_0_79])).
cnf(c_0_87, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_88, negated_conjecture, (~m1_subset_1(X1,u1_struct_0(esk2_0))|X1!=esk4_0|~m1_subset_1(X2,u1_struct_0(esk3_0))|X2!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_89, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk2_0))|v2_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_84]), c_0_85]), c_0_61]), c_0_54]), c_0_55])]), c_0_63]), c_0_62]), c_0_56]), c_0_57])).
cnf(c_0_90, negated_conjecture, (m1_pre_topc(k2_tsep_1(esk1_0,esk2_0,esk3_0),esk3_0)|~v2_pre_topc(X1)|~m1_pre_topc(k2_tsep_1(esk1_0,esk2_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_91, negated_conjecture, (~m1_subset_1(esk4_0,u1_struct_0(esk3_0))|~m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_88])])).
cnf(c_0_92, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_89]), c_0_54]), c_0_61]), c_0_55])]), c_0_56]), c_0_63]), c_0_57])).
cnf(c_0_93, negated_conjecture, (m1_pre_topc(k2_tsep_1(esk1_0,esk2_0,esk3_0),esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_41]), c_0_85]), c_0_54]), c_0_55]), c_0_61])]), c_0_56]), c_0_63]), c_0_57])).
cnf(c_0_94, negated_conjecture, (l1_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_54]), c_0_55])])).
cnf(c_0_95, negated_conjecture, (~m1_subset_1(esk4_0,u1_struct_0(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_92])])).
cnf(c_0_96, negated_conjecture, (v2_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_93]), c_0_94])]), c_0_95]), c_0_56])).
cnf(c_0_97, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_96]), c_0_54]), c_0_61]), c_0_55])]), c_0_56]), c_0_63]), c_0_57]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 98
# Proof object clause steps            : 59
# Proof object formula steps           : 39
# Proof object conjectures             : 29
# Proof object clause conjectures      : 26
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 30
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 81
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 8
# Initial clauses in saturation        : 25
# Processed clauses                    : 165
# ...of these trivial                  : 0
# ...subsumed                          : 37
# ...remaining for further processing  : 128
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 8
# Generated clauses                    : 448
# ...of the previous two non-trivial   : 419
# Contextual simplify-reflections      : 12
# Paramodulations                      : 445
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 87
#    Positive orientable unit clauses  : 19
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 61
# Current number of unprocessed clauses: 275
# ...number of literals in the above   : 4247
# Current number of archived formulas  : 0
# Current number of archived clauses   : 47
# Clause-clause subsumption calls (NU) : 2027
# Rec. Clause-clause subsumption calls : 307
# Non-unit clause-clause subsumptions  : 53
# Unit Clause-clause subsumption calls : 17
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 20360

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.003 s
# Total time               : 0.024 s
# Maximum resident set size: 28660 pages
