# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.006 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t113_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>(v1_funct_1(k7_tmap_1(X1,X2))&(v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))&(v5_pre_topc(k7_tmap_1(X1,X2),X1,k6_tmap_1(X1,X2))&m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2))))))))))), file('data/MPTP/proofs/10dcd5c501d845a0ae9554ad51ada133.E_input', t113_tmap_1)).
fof(t104_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(u1_struct_0(k6_tmap_1(X1,X2))=u1_struct_0(X1)&u1_pre_topc(k6_tmap_1(X1,X2))=k5_tmap_1(X1,X2)))), file('data/MPTP/proofs/10dcd5c501d845a0ae9554ad51ada133.E_input', t104_tmap_1)).
fof(dt_k6_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_pre_topc(k6_tmap_1(X1,X2))&(v2_pre_topc(k6_tmap_1(X1,X2))&l1_pre_topc(k6_tmap_1(X1,X2))))), file('data/MPTP/proofs/10dcd5c501d845a0ae9554ad51ada133.E_input', dt_k6_tmap_1)).
fof(rc2_tops_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>?[X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))&(v3_pre_topc(X2,X1)&v4_pre_topc(X2,X1)))), file('data/MPTP/proofs/10dcd5c501d845a0ae9554ad51ada133.E_input', rc2_tops_1)).
fof(d10_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k7_tmap_1(X1,X2)=k6_partfun1(u1_struct_0(X1)))), file('data/MPTP/proofs/10dcd5c501d845a0ae9554ad51ada133.E_input', d10_tmap_1)).
fof(t103_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X2,u1_pre_topc(X1))<=>u1_pre_topc(X1)=k5_tmap_1(X1,X2)))), file('data/MPTP/proofs/10dcd5c501d845a0ae9554ad51ada133.E_input', t103_tmap_1)).
fof(d5_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>r2_hidden(X2,u1_pre_topc(X1))))), file('data/MPTP/proofs/10dcd5c501d845a0ae9554ad51ada133.E_input', d5_pre_topc)).
fof(d4_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k3_struct_0(X1)=k6_partfun1(u1_struct_0(X1))), file('data/MPTP/proofs/10dcd5c501d845a0ae9554ad51ada133.E_input', d4_struct_0)).
fof(t171_funct_2, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k8_relset_1(X1,X1,k6_partfun1(X1),X2)=X2), file('data/MPTP/proofs/10dcd5c501d845a0ae9554ad51ada133.E_input', t171_funct_2)).
fof(dt_k7_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_funct_1(k7_tmap_1(X1,X2))&(v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))&m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))))), file('data/MPTP/proofs/10dcd5c501d845a0ae9554ad51ada133.E_input', dt_k7_tmap_1)).
fof(t55_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>((k2_struct_0(X2)=k1_xboole_0=>k2_struct_0(X1)=k1_xboole_0)=>(v5_pre_topc(X3,X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(v3_pre_topc(X4,X2)=>v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X1)))))))), file('data/MPTP/proofs/10dcd5c501d845a0ae9554ad51ada133.E_input', t55_tops_2)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/10dcd5c501d845a0ae9554ad51ada133.E_input', dt_l1_pre_topc)).
fof(redefinition_k8_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k8_relset_1(X1,X2,X3,X4)=k10_relat_1(X3,X4)), file('data/MPTP/proofs/10dcd5c501d845a0ae9554ad51ada133.E_input', redefinition_k8_relset_1)).
fof(fc3_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(v1_funct_1(k3_struct_0(X1))&(v1_funct_2(k3_struct_0(X1),u1_struct_0(X1),u1_struct_0(X1))&v5_pre_topc(k3_struct_0(X1),X1,X1)))), file('data/MPTP/proofs/10dcd5c501d845a0ae9554ad51ada133.E_input', fc3_tmap_1)).
fof(t105_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k6_tmap_1(X1,X2))))=>(X3=X2=>v3_pre_topc(X3,k6_tmap_1(X1,X2)))))), file('data/MPTP/proofs/10dcd5c501d845a0ae9554ad51ada133.E_input', t105_tmap_1)).
fof(dt_k3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>(v1_funct_1(k3_struct_0(X1))&(v1_funct_2(k3_struct_0(X1),u1_struct_0(X1),u1_struct_0(X1))&m1_subset_1(k3_struct_0(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))))), file('data/MPTP/proofs/10dcd5c501d845a0ae9554ad51ada133.E_input', dt_k3_struct_0)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/10dcd5c501d845a0ae9554ad51ada133.E_input', d3_struct_0)).
fof(c_0_17, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>(v1_funct_1(k7_tmap_1(X1,X2))&(v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))&(v5_pre_topc(k7_tmap_1(X1,X2),X1,k6_tmap_1(X1,X2))&m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))))))))), inference(assume_negation,[status(cth)],[t113_tmap_1])).
fof(c_0_18, plain, ![X7, X8]:((u1_struct_0(k6_tmap_1(X7,X8))=u1_struct_0(X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&(u1_pre_topc(k6_tmap_1(X7,X8))=k5_tmap_1(X7,X8)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t104_tmap_1])])])])])).
fof(c_0_19, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&((~v3_pre_topc(esk2_0,esk1_0)|(~v1_funct_1(k7_tmap_1(esk1_0,esk2_0))|(~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))|(~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))))))))&((v1_funct_1(k7_tmap_1(esk1_0,esk2_0))|v3_pre_topc(esk2_0,esk1_0))&((v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))|v3_pre_topc(esk2_0,esk1_0))&((v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|v3_pre_topc(esk2_0,esk1_0))&(m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))))|v3_pre_topc(esk2_0,esk1_0)))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])])).
fof(c_0_20, plain, ![X9, X10]:((v1_pre_topc(k6_tmap_1(X9,X10))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))))&((v2_pre_topc(k6_tmap_1(X9,X10))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))))&(l1_pre_topc(k6_tmap_1(X9,X10))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_tmap_1])])])])).
fof(c_0_21, plain, ![X11]:((m1_subset_1(esk3_1(X11),k1_zfmisc_1(u1_struct_0(X11)))|(~v2_pre_topc(X11)|~l1_pre_topc(X11)))&((v3_pre_topc(esk3_1(X11),X11)|(~v2_pre_topc(X11)|~l1_pre_topc(X11)))&(v4_pre_topc(esk3_1(X11),X11)|(~v2_pre_topc(X11)|~l1_pre_topc(X11))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[rc2_tops_1])])])])).
cnf(c_0_22, plain, (u1_struct_0(k6_tmap_1(X1,X2))=u1_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (l1_pre_topc(k6_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (v2_pre_topc(k6_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_29, plain, ![X13, X14]:(v2_struct_0(X13)|(~v2_pre_topc(X13)|~l1_pre_topc(X13))|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|k7_tmap_1(X13,X14)=k6_partfun1(u1_struct_0(X13)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_tmap_1])])])])).
cnf(c_0_30, plain, (m1_subset_1(esk3_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, negated_conjecture, (u1_struct_0(k6_tmap_1(esk1_0,esk2_0))=u1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_32, negated_conjecture, (l1_pre_topc(k6_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_23]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_33, negated_conjecture, (v2_pre_topc(k6_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_23]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_34, plain, (v2_struct_0(X1)|k7_tmap_1(X1,X2)=k6_partfun1(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_35, plain, ![X15, X16]:((~r2_hidden(X16,u1_pre_topc(X15))|u1_pre_topc(X15)=k5_tmap_1(X15,X16)|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X15)))|(v2_struct_0(X15)|(~v2_pre_topc(X15)|~l1_pre_topc(X15))))&(u1_pre_topc(X15)!=k5_tmap_1(X15,X16)|r2_hidden(X16,u1_pre_topc(X15))|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X15)))|(v2_struct_0(X15)|(~v2_pre_topc(X15)|~l1_pre_topc(X15))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t103_tmap_1])])])])])).
fof(c_0_36, plain, ![X17, X18]:((~v3_pre_topc(X18,X17)|r2_hidden(X18,u1_pre_topc(X17))|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))|~l1_pre_topc(X17))&(~r2_hidden(X18,u1_pre_topc(X17))|v3_pre_topc(X18,X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))|~l1_pre_topc(X17))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_pre_topc])])])])).
fof(c_0_37, plain, ![X19]:(~l1_struct_0(X19)|k3_struct_0(X19)=k6_partfun1(u1_struct_0(X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_struct_0])])).
fof(c_0_38, plain, ![X20, X21]:(~m1_subset_1(X21,k1_zfmisc_1(X20))|k8_relset_1(X20,X20,k6_partfun1(X20),X21)=X21), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t171_funct_2])])).
fof(c_0_39, plain, ![X22, X23]:((v1_funct_1(k7_tmap_1(X22,X23))|(v2_struct_0(X22)|(~v2_pre_topc(X22)|~l1_pre_topc(X22))|~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(X22)))))&((v1_funct_2(k7_tmap_1(X22,X23),u1_struct_0(X22),u1_struct_0(k6_tmap_1(X22,X23)))|(v2_struct_0(X22)|(~v2_pre_topc(X22)|~l1_pre_topc(X22))|~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(X22)))))&(m1_subset_1(k7_tmap_1(X22,X23),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(k6_tmap_1(X22,X23)))))|(v2_struct_0(X22)|(~v2_pre_topc(X22)|~l1_pre_topc(X22))|~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(X22))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_tmap_1])])])])).
cnf(c_0_40, negated_conjecture, (m1_subset_1(esk3_1(k6_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32]), c_0_33])])).
cnf(c_0_41, negated_conjecture, (k6_partfun1(u1_struct_0(esk1_0))=k7_tmap_1(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_23]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_42, plain, (u1_pre_topc(k6_tmap_1(X1,X2))=k5_tmap_1(X1,X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_43, plain, (u1_pre_topc(X2)=k5_tmap_1(X2,X1)|v2_struct_0(X2)|~r2_hidden(X1,u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_44, plain, (r2_hidden(X1,u1_pre_topc(X2))|~v3_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_45, plain, ![X24, X25, X26, X27]:(((~v5_pre_topc(X26,X24,X25)|(~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X25)))|(~v3_pre_topc(X27,X25)|v3_pre_topc(k8_relset_1(u1_struct_0(X24),u1_struct_0(X25),X26,X27),X24)))|k2_struct_0(X25)=k1_xboole_0|(~v1_funct_1(X26)|(~v1_funct_2(X26,u1_struct_0(X24),u1_struct_0(X25))|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X24),u1_struct_0(X25))))))|~l1_pre_topc(X25)|~l1_pre_topc(X24))&((m1_subset_1(esk4_3(X24,X25,X26),k1_zfmisc_1(u1_struct_0(X25)))|v5_pre_topc(X26,X24,X25)|k2_struct_0(X25)=k1_xboole_0|(~v1_funct_1(X26)|(~v1_funct_2(X26,u1_struct_0(X24),u1_struct_0(X25))|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X24),u1_struct_0(X25))))))|~l1_pre_topc(X25)|~l1_pre_topc(X24))&((v3_pre_topc(esk4_3(X24,X25,X26),X25)|v5_pre_topc(X26,X24,X25)|k2_struct_0(X25)=k1_xboole_0|(~v1_funct_1(X26)|(~v1_funct_2(X26,u1_struct_0(X24),u1_struct_0(X25))|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X24),u1_struct_0(X25))))))|~l1_pre_topc(X25)|~l1_pre_topc(X24))&(~v3_pre_topc(k8_relset_1(u1_struct_0(X24),u1_struct_0(X25),X26,esk4_3(X24,X25,X26)),X24)|v5_pre_topc(X26,X24,X25)|k2_struct_0(X25)=k1_xboole_0|(~v1_funct_1(X26)|(~v1_funct_2(X26,u1_struct_0(X24),u1_struct_0(X25))|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X24),u1_struct_0(X25))))))|~l1_pre_topc(X25)|~l1_pre_topc(X24)))))&((~v5_pre_topc(X26,X24,X25)|(~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X25)))|(~v3_pre_topc(X27,X25)|v3_pre_topc(k8_relset_1(u1_struct_0(X24),u1_struct_0(X25),X26,X27),X24)))|k2_struct_0(X24)!=k1_xboole_0|(~v1_funct_1(X26)|(~v1_funct_2(X26,u1_struct_0(X24),u1_struct_0(X25))|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X24),u1_struct_0(X25))))))|~l1_pre_topc(X25)|~l1_pre_topc(X24))&((m1_subset_1(esk4_3(X24,X25,X26),k1_zfmisc_1(u1_struct_0(X25)))|v5_pre_topc(X26,X24,X25)|k2_struct_0(X24)!=k1_xboole_0|(~v1_funct_1(X26)|(~v1_funct_2(X26,u1_struct_0(X24),u1_struct_0(X25))|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X24),u1_struct_0(X25))))))|~l1_pre_topc(X25)|~l1_pre_topc(X24))&((v3_pre_topc(esk4_3(X24,X25,X26),X25)|v5_pre_topc(X26,X24,X25)|k2_struct_0(X24)!=k1_xboole_0|(~v1_funct_1(X26)|(~v1_funct_2(X26,u1_struct_0(X24),u1_struct_0(X25))|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X24),u1_struct_0(X25))))))|~l1_pre_topc(X25)|~l1_pre_topc(X24))&(~v3_pre_topc(k8_relset_1(u1_struct_0(X24),u1_struct_0(X25),X26,esk4_3(X24,X25,X26)),X24)|v5_pre_topc(X26,X24,X25)|k2_struct_0(X24)!=k1_xboole_0|(~v1_funct_1(X26)|(~v1_funct_2(X26,u1_struct_0(X24),u1_struct_0(X25))|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X24),u1_struct_0(X25))))))|~l1_pre_topc(X25)|~l1_pre_topc(X24)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_tops_2])])])])])).
cnf(c_0_46, plain, (k3_struct_0(X1)=k6_partfun1(u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_47, plain, ![X29]:(~l1_pre_topc(X29)|l1_struct_0(X29)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_48, plain, ![X30, X31, X32, X33]:(~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(X30,X31)))|k8_relset_1(X30,X31,X32,X33)=k10_relat_1(X32,X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_relset_1])])).
cnf(c_0_49, plain, (k8_relset_1(X2,X2,k6_partfun1(X2),X1)=X1|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, plain, (m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51, negated_conjecture, (k7_tmap_1(esk1_0,esk3_1(k6_tmap_1(esk1_0,esk2_0)))=k7_tmap_1(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_40]), c_0_24]), c_0_25])]), c_0_26]), c_0_41])).
cnf(c_0_52, negated_conjecture, (u1_struct_0(k6_tmap_1(esk1_0,esk3_1(k6_tmap_1(esk1_0,esk2_0))))=u1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_40]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_53, negated_conjecture, (u1_pre_topc(k6_tmap_1(esk1_0,esk2_0))=k5_tmap_1(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_23]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_54, plain, (k5_tmap_1(X1,X2)=u1_pre_topc(X1)|v2_struct_0(X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_55, plain, (v5_pre_topc(X3,X1,X2)|k2_struct_0(X2)=k1_xboole_0|~v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk4_3(X1,X2,X3)),X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_56, plain, ![X34]:((v1_funct_1(k3_struct_0(X34))|(v2_struct_0(X34)|(~v2_pre_topc(X34)|~l1_pre_topc(X34))))&((v1_funct_2(k3_struct_0(X34),u1_struct_0(X34),u1_struct_0(X34))|(v2_struct_0(X34)|(~v2_pre_topc(X34)|~l1_pre_topc(X34))))&(v5_pre_topc(k3_struct_0(X34),X34,X34)|(v2_struct_0(X34)|(~v2_pre_topc(X34)|~l1_pre_topc(X34)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc3_tmap_1])])])])).
cnf(c_0_57, negated_conjecture, (k3_struct_0(esk1_0)=k7_tmap_1(esk1_0,esk2_0)|~l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_46, c_0_41])).
cnf(c_0_58, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_59, plain, (v1_funct_1(k7_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_60, plain, (k8_relset_1(X2,X3,X1,X4)=k10_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_61, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,esk2_0),X1)=X1|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_49, c_0_41])).
cnf(c_0_62, negated_conjecture, (m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_40]), c_0_51]), c_0_52]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_63, negated_conjecture, (r2_hidden(X1,k5_tmap_1(esk1_0,esk2_0))|~v3_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_53]), c_0_31]), c_0_32])])).
cnf(c_0_64, negated_conjecture, (k5_tmap_1(esk1_0,esk2_0)=u1_pre_topc(esk1_0)|~v3_pre_topc(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_23]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_65, plain, (v3_pre_topc(esk4_3(X1,X2,X3),X2)|v5_pre_topc(X3,X1,X2)|k2_struct_0(X2)=k1_xboole_0|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_66, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v5_pre_topc(X1,X2,k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~v3_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(esk1_0),X1,esk4_3(X2,k6_tmap_1(esk1_0,esk2_0),X1)),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_31]), c_0_32])])).
cnf(c_0_67, plain, (v1_funct_2(k3_struct_0(X1),u1_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_68, negated_conjecture, (k3_struct_0(esk1_0)=k7_tmap_1(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_24])])).
cnf(c_0_69, negated_conjecture, (v1_funct_1(k7_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_23]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_70, negated_conjecture, (k10_relat_1(k7_tmap_1(esk1_0,esk2_0),X1)=X1|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62])])).
fof(c_0_71, plain, ![X35, X36, X37]:(v2_struct_0(X35)|(~v2_pre_topc(X35)|~l1_pre_topc(X35))|(~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X35)))|(~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(k6_tmap_1(X35,X36))))|(X37!=X36|v3_pre_topc(X37,k6_tmap_1(X35,X36)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t105_tmap_1])])])])).
cnf(c_0_72, plain, (v3_pre_topc(X1,X2)|~r2_hidden(X1,u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_73, negated_conjecture, (r2_hidden(X1,u1_pre_topc(esk1_0))|~v3_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0))|~v3_pre_topc(esk2_0,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_74, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v5_pre_topc(X1,X2,k6_tmap_1(esk1_0,esk2_0))|v3_pre_topc(esk4_3(X2,k6_tmap_1(esk1_0,esk2_0),X1),k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_31]), c_0_32])])).
cnf(c_0_75, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~v3_pre_topc(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,esk2_0),esk4_3(esk1_0,k6_tmap_1(esk1_0,esk2_0),k7_tmap_1(esk1_0,esk2_0))),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68]), c_0_68]), c_0_69]), c_0_68]), c_0_68]), c_0_68]), c_0_62]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_76, negated_conjecture, (k8_relset_1(X1,X2,k7_tmap_1(esk1_0,esk2_0),X3)=X3|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_60, c_0_70])).
cnf(c_0_77, negated_conjecture, (~v3_pre_topc(esk2_0,esk1_0)|~v1_funct_1(k7_tmap_1(esk1_0,esk2_0))|~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))|~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_78, plain, (v2_struct_0(X1)|v3_pre_topc(X3,k6_tmap_1(X1,X2))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k6_tmap_1(X1,X2))))|X3!=X2), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_79, plain, (m1_subset_1(esk4_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|v5_pre_topc(X3,X1,X2)|k2_struct_0(X2)=k1_xboole_0|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_80, plain, ![X38]:((v1_funct_1(k3_struct_0(X38))|~l1_struct_0(X38))&((v1_funct_2(k3_struct_0(X38),u1_struct_0(X38),u1_struct_0(X38))|~l1_struct_0(X38))&(m1_subset_1(k3_struct_0(X38),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X38),u1_struct_0(X38))))|~l1_struct_0(X38)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_struct_0])])])).
cnf(c_0_81, negated_conjecture, (v3_pre_topc(X1,esk1_0)|~v3_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0))|~v3_pre_topc(esk2_0,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_24])])).
cnf(c_0_82, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|v3_pre_topc(esk4_3(esk1_0,k6_tmap_1(esk1_0,esk2_0),k7_tmap_1(esk1_0,esk2_0)),k6_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_67]), c_0_68]), c_0_68]), c_0_68]), c_0_69]), c_0_68]), c_0_62]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_83, negated_conjecture, (v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|v3_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_84, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~v3_pre_topc(esk4_3(esk1_0,k6_tmap_1(esk1_0,esk2_0),k7_tmap_1(esk1_0,esk2_0)),esk1_0)|~m1_subset_1(esk4_3(esk1_0,k6_tmap_1(esk1_0,esk2_0),k7_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_62])])).
cnf(c_0_85, negated_conjecture, (~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))|~v3_pre_topc(esk2_0,esk1_0)|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_69])])).
cnf(c_0_86, plain, (k10_relat_1(k6_partfun1(X1),X2)=X2|~m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_49, c_0_60])).
cnf(c_0_87, plain, (v3_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(X3),X1,X4),X2)|k2_struct_0(X3)=k1_xboole_0|~v5_pre_topc(X1,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~v3_pre_topc(X4,X3)|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(X3)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_88, plain, (v3_pre_topc(X1,k6_tmap_1(X2,X1))|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k6_tmap_1(X2,X1))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(er,[status(thm)],[c_0_78])).
cnf(c_0_89, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v5_pre_topc(X1,X2,k6_tmap_1(esk1_0,esk2_0))|m1_subset_1(esk4_3(X2,k6_tmap_1(esk1_0,esk2_0),X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_31]), c_0_32])])).
cnf(c_0_90, plain, (v1_funct_2(k3_struct_0(X1),u1_struct_0(X1),u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_91, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(esk4_3(esk1_0,k6_tmap_1(esk1_0,esk2_0),k7_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_83]), c_0_84])).
cnf(c_0_92, negated_conjecture, (~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v3_pre_topc(esk2_0,esk1_0)|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_31]), c_0_31])).
cnf(c_0_93, negated_conjecture, (v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_94, negated_conjecture, (k10_relat_1(k7_tmap_1(esk1_0,esk2_0),esk2_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_23]), c_0_41]), c_0_41]), c_0_62])])).
cnf(c_0_95, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(esk1_0),X2,X3),X1)|~v5_pre_topc(X2,X1,k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2)|~v3_pre_topc(X3,k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_31]), c_0_32])])).
cnf(c_0_96, negated_conjecture, (v3_pre_topc(esk2_0,k6_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_23]), c_0_31]), c_0_23]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_97, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~l1_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_68]), c_0_68]), c_0_68]), c_0_69]), c_0_68]), c_0_62]), c_0_24])]), c_0_91])).
cnf(c_0_98, negated_conjecture, (~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~v3_pre_topc(esk2_0,esk1_0)|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_93])])).
cnf(c_0_99, plain, (v3_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(X3),X1,X4),X2)|~v5_pre_topc(X1,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~v3_pre_topc(X4,X3)|k2_struct_0(X2)!=k1_xboole_0|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(X3)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_100, negated_conjecture, (k8_relset_1(X1,X2,k7_tmap_1(esk1_0,esk2_0),esk2_0)=esk2_0|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_60, c_0_94])).
fof(c_0_101, plain, ![X39]:(~l1_struct_0(X39)|k2_struct_0(X39)=u1_struct_0(X39)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_102, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(esk1_0),X2,esk2_0),X1)|~v5_pre_topc(X2,X1,k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_23]), c_0_96])])).
cnf(c_0_103, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_58]), c_0_24])])).
cnf(c_0_104, negated_conjecture, (~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~v3_pre_topc(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_98, c_0_62])])).
cnf(c_0_105, negated_conjecture, (v3_pre_topc(esk2_0,X1)|k2_struct_0(X1)!=k1_xboole_0|~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),X1,X2)|~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(X1),u1_struct_0(X2))|~v3_pre_topc(esk2_0,X2)|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_69])])).
cnf(c_0_106, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_107, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v3_pre_topc(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,esk2_0),esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_93]), c_0_69]), c_0_62]), c_0_24])])).
cnf(c_0_108, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|~v3_pre_topc(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_104, c_0_103])).
cnf(c_0_109, negated_conjecture, (v3_pre_topc(esk2_0,X1)|k2_struct_0(X1)!=k1_xboole_0|~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),X1,k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(X1),u1_struct_0(esk1_0))|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_31]), c_0_96]), c_0_23]), c_0_32])])).
cnf(c_0_110, plain, (v3_pre_topc(esk4_3(X1,X2,X3),X2)|v5_pre_topc(X3,X1,X2)|k2_struct_0(X1)!=k1_xboole_0|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_111, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_106, c_0_58])).
cnf(c_0_112, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_60]), c_0_94]), c_0_62])]), c_0_108])).
cnf(c_0_113, negated_conjecture, (v3_pre_topc(esk2_0,esk1_0)|k2_struct_0(esk1_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_83]), c_0_93]), c_0_62]), c_0_24])])).
cnf(c_0_114, plain, (v5_pre_topc(X1,X2,X3)|v3_pre_topc(esk4_3(X2,X3,X1),X3)|u1_struct_0(X2)!=k1_xboole_0|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(X3)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_110, c_0_111])).
cnf(c_0_115, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_31]), c_0_32])])).
cnf(c_0_116, negated_conjecture, (v3_pre_topc(esk2_0,esk1_0)|u1_struct_0(esk1_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_111]), c_0_24])])).
cnf(c_0_117, plain, (v5_pre_topc(X3,X1,X2)|~v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk4_3(X1,X2,X3)),X1)|k2_struct_0(X1)!=k1_xboole_0|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_118, negated_conjecture, (v5_pre_topc(X1,esk1_0,X2)|v3_pre_topc(esk4_3(esk1_0,X2,X1),X2)|~v1_funct_2(X1,k1_xboole_0,u1_struct_0(X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,u1_struct_0(X2))))|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_24])])).
cnf(c_0_119, negated_conjecture, (u1_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_31, c_0_115])).
cnf(c_0_120, negated_conjecture, (v3_pre_topc(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_116, c_0_115])])).
cnf(c_0_121, plain, (m1_subset_1(esk4_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|v5_pre_topc(X3,X1,X2)|k2_struct_0(X1)!=k1_xboole_0|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_122, negated_conjecture, (v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),X1,X2)|k2_struct_0(X1)!=k1_xboole_0|~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(X1),u1_struct_0(X2))|~v3_pre_topc(esk4_3(X1,X2,k7_tmap_1(esk1_0,esk2_0)),X1)|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk4_3(X1,X2,k7_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_76]), c_0_69])])).
cnf(c_0_123, negated_conjecture, (v5_pre_topc(X1,esk1_0,k6_tmap_1(esk1_0,esk2_0))|v3_pre_topc(esk4_3(esk1_0,k6_tmap_1(esk1_0,esk2_0),X1),k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(X1,k1_xboole_0,k1_xboole_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_32])])).
cnf(c_0_124, negated_conjecture, (v1_funct_2(k7_tmap_1(esk1_0,esk2_0),k1_xboole_0,k1_xboole_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_115]), c_0_115])).
cnf(c_0_125, negated_conjecture, (m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_115]), c_0_115])).
cnf(c_0_126, negated_conjecture, (~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_104, c_0_120])])).
cnf(c_0_127, plain, (v5_pre_topc(X1,X2,X3)|m1_subset_1(esk4_3(X2,X3,X1),k1_zfmisc_1(u1_struct_0(X3)))|u1_struct_0(X2)!=k1_xboole_0|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(X3)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_121, c_0_111])).
cnf(c_0_128, negated_conjecture, (v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),X1,X2)|u1_struct_0(X1)!=k1_xboole_0|~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(X1),u1_struct_0(X2))|~v3_pre_topc(esk4_3(X1,X2,k7_tmap_1(esk1_0,esk2_0)),X1)|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk4_3(X1,X2,k7_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_122, c_0_111])).
cnf(c_0_129, negated_conjecture, (v3_pre_topc(X1,esk1_0)|~v3_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_115]), c_0_120])])).
cnf(c_0_130, negated_conjecture, (v3_pre_topc(esk4_3(esk1_0,k6_tmap_1(esk1_0,esk2_0),k7_tmap_1(esk1_0,esk2_0)),k6_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124]), c_0_69]), c_0_125])]), c_0_126])).
cnf(c_0_131, negated_conjecture, (v5_pre_topc(X1,esk1_0,X2)|m1_subset_1(esk4_3(esk1_0,X2,X1),k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X1,k1_xboole_0,u1_struct_0(X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,u1_struct_0(X2))))|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_115]), c_0_24])])).
cnf(c_0_132, negated_conjecture, (v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),X1,X2)|u1_struct_0(X1)!=k1_xboole_0|~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(X1),u1_struct_0(X2))|~v3_pre_topc(esk4_3(X1,X2,k7_tmap_1(esk1_0,esk2_0)),X1)|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk4_3(X1,X2,k7_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(k1_xboole_0))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(rw,[status(thm)],[c_0_128, c_0_115])).
cnf(c_0_133, negated_conjecture, (v3_pre_topc(esk4_3(esk1_0,k6_tmap_1(esk1_0,esk2_0),k7_tmap_1(esk1_0,esk2_0)),esk1_0)|~m1_subset_1(esk4_3(esk1_0,k6_tmap_1(esk1_0,esk2_0),k7_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_129, c_0_130])).
cnf(c_0_134, negated_conjecture, (v5_pre_topc(X1,esk1_0,k6_tmap_1(esk1_0,esk2_0))|m1_subset_1(esk4_3(esk1_0,k6_tmap_1(esk1_0,esk2_0),X1),k1_zfmisc_1(k1_xboole_0))|~v1_funct_2(X1,k1_xboole_0,k1_xboole_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_119]), c_0_32])])).
cnf(c_0_135, negated_conjecture, (~m1_subset_1(esk4_3(esk1_0,k6_tmap_1(esk1_0,esk2_0),k7_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(k1_xboole_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_133]), c_0_115]), c_0_115]), c_0_31]), c_0_115]), c_0_124]), c_0_115]), c_0_31]), c_0_115]), c_0_125]), c_0_32]), c_0_24])]), c_0_126])).
cnf(c_0_136, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_124]), c_0_69]), c_0_125])]), c_0_126]), c_0_135]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 137
# Proof object clause steps            : 102
# Proof object formula steps           : 35
# Proof object conjectures             : 72
# Proof object clause conjectures      : 69
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 33
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 57
# Proof object simplifying inferences  : 193
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 45
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 45
# Processed clauses                    : 4421
# ...of these trivial                  : 79
# ...subsumed                          : 2292
# ...remaining for further processing  : 2050
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 157
# Backward-rewritten                   : 506
# Generated clauses                    : 23737
# ...of the previous two non-trivial   : 22611
# Contextual simplify-reflections      : 354
# Paramodulations                      : 23736
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1341
#    Positive orientable unit clauses  : 92
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 1246
# Current number of unprocessed clauses: 18175
# ...number of literals in the above   : 121385
# Current number of archived formulas  : 0
# Current number of archived clauses   : 708
# Clause-clause subsumption calls (NU) : 533533
# Rec. Clause-clause subsumption calls : 129609
# Non-unit clause-clause subsumptions  : 2756
# Unit Clause-clause subsumption calls : 8942
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 334
# BW rewrite match successes           : 32
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1391980

# -------------------------------------------------
# User time                : 0.418 s
# System time              : 0.007 s
# Total time               : 0.425 s
# Maximum resident set size: 28788 pages
