# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t116_funct_2, conjecture, ![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X5]:~((r2_hidden(X5,k7_relset_1(X1,X2,X4,X3))&![X6]:(m1_subset_1(X6,X1)=>~((r2_hidden(X6,X3)&X5=k1_funct_1(X4,X6))))))), file('data/MPTP/proofs/10f91ff0a4fd48c8abb6d1541804a83a.E_input', t116_funct_2)).
fof(redefinition_k7_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k7_relset_1(X1,X2,X3,X4)=k9_relat_1(X3,X4)), file('data/MPTP/proofs/10f91ff0a4fd48c8abb6d1541804a83a.E_input', redefinition_k7_relset_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/10f91ff0a4fd48c8abb6d1541804a83a.E_input', d2_subset_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/10f91ff0a4fd48c8abb6d1541804a83a.E_input', d1_xboole_0)).
fof(t143_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k9_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k1_relat_1(X3))&(r2_hidden(k4_tarski(X4,X1),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/10f91ff0a4fd48c8abb6d1541804a83a.E_input', t143_relat_1)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/10f91ff0a4fd48c8abb6d1541804a83a.E_input', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/10f91ff0a4fd48c8abb6d1541804a83a.E_input', fc6_relat_1)).
fof(dt_k7_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k7_relset_1(X1,X2,X3,X4),k1_zfmisc_1(X2))), file('data/MPTP/proofs/10f91ff0a4fd48c8abb6d1541804a83a.E_input', dt_k7_relset_1)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/10f91ff0a4fd48c8abb6d1541804a83a.E_input', t4_subset)).
fof(t8_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>(r2_hidden(X1,k1_relat_1(X3))&X2=k1_funct_1(X3,X1)))), file('data/MPTP/proofs/10f91ff0a4fd48c8abb6d1541804a83a.E_input', t8_funct_1)).
fof(t52_relset_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(~(v1_xboole_0(X3))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))=>![X5]:(m1_subset_1(X5,X1)=>(r2_hidden(X5,k7_relset_1(X3,X1,X4,X2))<=>?[X6]:(m1_subset_1(X6,X3)&(r2_hidden(k4_tarski(X6,X5),X4)&r2_hidden(X6,X2))))))))), file('data/MPTP/proofs/10f91ff0a4fd48c8abb6d1541804a83a.E_input', t52_relset_1)).
fof(cc3_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/10f91ff0a4fd48c8abb6d1541804a83a.E_input', cc3_relset_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/10f91ff0a4fd48c8abb6d1541804a83a.E_input', t1_subset)).
fof(cc4_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/10f91ff0a4fd48c8abb6d1541804a83a.E_input', cc4_relset_1)).
fof(c_0_14, negated_conjecture, ~(![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X5]:~((r2_hidden(X5,k7_relset_1(X1,X2,X4,X3))&![X6]:(m1_subset_1(X6,X1)=>~((r2_hidden(X6,X3)&X5=k1_funct_1(X4,X6)))))))), inference(assume_negation,[status(cth)],[t116_funct_2])).
fof(c_0_15, plain, ![X53, X54, X55, X56]:(~m1_subset_1(X55,k1_zfmisc_1(k2_zfmisc_1(X53,X54)))|k7_relset_1(X53,X54,X55,X56)=k9_relat_1(X55,X56)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_relset_1])])).
fof(c_0_16, negated_conjecture, ![X12]:((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,esk1_0,esk2_0)&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))&(r2_hidden(esk5_0,k7_relset_1(esk1_0,esk2_0,esk4_0,esk3_0))&(~m1_subset_1(X12,esk1_0)|(~r2_hidden(X12,esk3_0)|esk5_0!=k1_funct_1(esk4_0,X12))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])])).
fof(c_0_17, plain, ![X22, X23]:(((~m1_subset_1(X23,X22)|r2_hidden(X23,X22)|v1_xboole_0(X22))&(~r2_hidden(X23,X22)|m1_subset_1(X23,X22)|v1_xboole_0(X22)))&((~m1_subset_1(X23,X22)|v1_xboole_0(X23)|~v1_xboole_0(X22))&(~v1_xboole_0(X23)|m1_subset_1(X23,X22)|~v1_xboole_0(X22)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_18, plain, ![X30, X31, X32]:((~v1_xboole_0(X30)|~r2_hidden(X31,X30))&(r2_hidden(esk7_1(X32),X32)|v1_xboole_0(X32))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_19, plain, ![X44, X45, X46, X48]:(((r2_hidden(esk8_3(X44,X45,X46),k1_relat_1(X46))|~r2_hidden(X44,k9_relat_1(X46,X45))|~v1_relat_1(X46))&((r2_hidden(k4_tarski(esk8_3(X44,X45,X46),X44),X46)|~r2_hidden(X44,k9_relat_1(X46,X45))|~v1_relat_1(X46))&(r2_hidden(esk8_3(X44,X45,X46),X45)|~r2_hidden(X44,k9_relat_1(X46,X45))|~v1_relat_1(X46))))&(~r2_hidden(X48,k1_relat_1(X46))|(~r2_hidden(k4_tarski(X48,X44),X46)|~r2_hidden(X48,X45))|r2_hidden(X44,k9_relat_1(X46,X45))|~v1_relat_1(X46))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t143_relat_1])])])])])).
cnf(c_0_20, plain, (k7_relset_1(X2,X3,X1,X4)=k9_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, plain, ![X13, X14]:(~v1_relat_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(X13))|v1_relat_1(X14))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_23, plain, ![X34, X35]:v1_relat_1(k2_zfmisc_1(X34,X35)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_24, plain, ![X49, X50, X51, X52]:(~m1_subset_1(X51,k1_zfmisc_1(k2_zfmisc_1(X49,X50)))|m1_subset_1(k7_relset_1(X49,X50,X51,X52),k1_zfmisc_1(X50))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_relset_1])])).
fof(c_0_25, plain, ![X24, X25, X26]:(~r2_hidden(X24,X25)|~m1_subset_1(X25,k1_zfmisc_1(X26))|m1_subset_1(X24,X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_26, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_29, plain, (r2_hidden(k4_tarski(esk8_3(X1,X2,X3),X1),X3)|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_30, negated_conjecture, (r2_hidden(esk5_0,k7_relset_1(esk1_0,esk2_0,esk4_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_31, negated_conjecture, (k7_relset_1(esk1_0,esk2_0,esk4_0,X1)=k9_relat_1(esk4_0,X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_32, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_34, plain, (m1_subset_1(k7_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_35, plain, ![X27, X28, X29]:(((r2_hidden(X27,k1_relat_1(X29))|~r2_hidden(k4_tarski(X27,X28),X29)|(~v1_relat_1(X29)|~v1_funct_1(X29)))&(X28=k1_funct_1(X29,X27)|~r2_hidden(k4_tarski(X27,X28),X29)|(~v1_relat_1(X29)|~v1_funct_1(X29))))&(~r2_hidden(X27,k1_relat_1(X29))|X28!=k1_funct_1(X29,X27)|r2_hidden(k4_tarski(X27,X28),X29)|(~v1_relat_1(X29)|~v1_funct_1(X29)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_funct_1])])])).
fof(c_0_36, plain, ![X15, X16, X17, X18, X19, X21]:(((m1_subset_1(esk6_5(X15,X16,X17,X18,X19),X17)|~r2_hidden(X19,k7_relset_1(X17,X15,X18,X16))|~m1_subset_1(X19,X15)|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(X17,X15)))|v1_xboole_0(X17)|v1_xboole_0(X16)|v1_xboole_0(X15))&((r2_hidden(k4_tarski(esk6_5(X15,X16,X17,X18,X19),X19),X18)|~r2_hidden(X19,k7_relset_1(X17,X15,X18,X16))|~m1_subset_1(X19,X15)|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(X17,X15)))|v1_xboole_0(X17)|v1_xboole_0(X16)|v1_xboole_0(X15))&(r2_hidden(esk6_5(X15,X16,X17,X18,X19),X16)|~r2_hidden(X19,k7_relset_1(X17,X15,X18,X16))|~m1_subset_1(X19,X15)|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(X17,X15)))|v1_xboole_0(X17)|v1_xboole_0(X16)|v1_xboole_0(X15))))&(~m1_subset_1(X21,X17)|(~r2_hidden(k4_tarski(X21,X19),X18)|~r2_hidden(X21,X16))|r2_hidden(X19,k7_relset_1(X17,X15,X18,X16))|~m1_subset_1(X19,X15)|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(X17,X15)))|v1_xboole_0(X17)|v1_xboole_0(X16)|v1_xboole_0(X15))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t52_relset_1])])])])])])).
cnf(c_0_37, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_38, plain, ![X36, X37, X38]:(~v1_xboole_0(X36)|(~m1_subset_1(X38,k1_zfmisc_1(k2_zfmisc_1(X36,X37)))|v1_xboole_0(X38))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_relset_1])])])).
fof(c_0_39, plain, ![X39, X40]:(~r2_hidden(X39,X40)|m1_subset_1(X39,X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_40, negated_conjecture, (v1_xboole_0(esk4_0)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_26, c_0_21])).
cnf(c_0_41, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))|r2_hidden(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_27, c_0_21])).
cnf(c_0_42, plain, (~v1_xboole_0(X1)|~v1_relat_1(X1)|~r2_hidden(X2,k9_relat_1(X1,X3))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_43, negated_conjecture, (r2_hidden(esk5_0,k9_relat_1(esk4_0,esk3_0))), inference(rw,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_44, negated_conjecture, (v1_relat_1(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_21]), c_0_33])])).
cnf(c_0_45, negated_conjecture, (m1_subset_1(k9_relat_1(esk4_0,X1),k1_zfmisc_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_31]), c_0_21])])).
cnf(c_0_46, plain, (X1=k1_funct_1(X2,X3)|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_47, plain, (r2_hidden(k4_tarski(esk6_5(X1,X2,X3,X4,X5),X5),X4)|v1_xboole_0(X3)|v1_xboole_0(X2)|v1_xboole_0(X1)|~r2_hidden(X5,k7_relset_1(X3,X1,X4,X2))|~m1_subset_1(X5,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_48, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relset_1(X3,X2,X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(spm,[status(thm)],[c_0_37, c_0_34])).
cnf(c_0_49, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51, negated_conjecture, (v1_xboole_0(esk4_0)|r2_hidden(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_52, negated_conjecture, (~v1_xboole_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])])).
fof(c_0_53, plain, ![X41, X42, X43]:(~v1_xboole_0(X41)|(~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(X42,X41)))|v1_xboole_0(X43))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(X1,esk2_0)|~r2_hidden(X1,k9_relat_1(esk4_0,X2))), inference(spm,[status(thm)],[c_0_37, c_0_45])).
cnf(c_0_55, negated_conjecture, (~m1_subset_1(X1,esk1_0)|~r2_hidden(X1,esk3_0)|esk5_0!=k1_funct_1(esk4_0,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_56, plain, (k1_funct_1(X1,esk6_5(X2,X3,X4,X1,X5))=X5|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X4)|~v1_relat_1(X1)|~r2_hidden(X5,k7_relset_1(X4,X2,X1,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])).
cnf(c_0_57, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_58, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_59, negated_conjecture, (r2_hidden(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(sr,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_60, plain, (r2_hidden(esk8_3(X1,X2,X3),X2)|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_61, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_62, negated_conjecture, (m1_subset_1(esk5_0,esk2_0)), inference(spm,[status(thm)],[c_0_54, c_0_43])).
cnf(c_0_63, negated_conjecture, (v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~r2_hidden(esk6_5(X3,X2,X1,esk4_0,esk5_0),esk3_0)|~r2_hidden(esk5_0,k7_relset_1(X1,X3,esk4_0,X2))|~m1_subset_1(esk6_5(X3,X2,X1,esk4_0,esk5_0),esk1_0)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_44]), c_0_57])])])).
cnf(c_0_64, plain, (m1_subset_1(esk6_5(X1,X2,X3,X4,X5),X3)|v1_xboole_0(X3)|v1_xboole_0(X2)|v1_xboole_0(X1)|~r2_hidden(X5,k7_relset_1(X3,X1,X4,X2))|~m1_subset_1(X5,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_65, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_52])).
cnf(c_0_66, plain, (~v1_xboole_0(X1)|~v1_relat_1(X2)|~r2_hidden(X3,k9_relat_1(X2,X1))), inference(spm,[status(thm)],[c_0_28, c_0_60])).
cnf(c_0_67, negated_conjecture, (v1_xboole_0(esk4_0)|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_61, c_0_21])).
cnf(c_0_68, negated_conjecture, (v1_xboole_0(esk2_0)|r2_hidden(esk5_0,esk2_0)), inference(spm,[status(thm)],[c_0_27, c_0_62])).
cnf(c_0_69, negated_conjecture, (v1_xboole_0(X1)|v1_xboole_0(X2)|~r2_hidden(esk6_5(X1,X2,esk1_0,esk4_0,esk5_0),esk3_0)|~r2_hidden(esk5_0,k7_relset_1(esk1_0,X1,esk4_0,X2))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,X1)))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_65]), c_0_48])).
cnf(c_0_70, plain, (r2_hidden(esk6_5(X1,X2,X3,X4,X5),X2)|v1_xboole_0(X3)|v1_xboole_0(X2)|v1_xboole_0(X1)|~r2_hidden(X5,k7_relset_1(X3,X1,X4,X2))|~m1_subset_1(X5,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_71, negated_conjecture, (~v1_xboole_0(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_43]), c_0_44])])).
cnf(c_0_72, negated_conjecture, (r2_hidden(esk5_0,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_52])).
cnf(c_0_73, negated_conjecture, (v1_xboole_0(X1)|~r2_hidden(esk5_0,k7_relset_1(esk1_0,X1,esk4_0,esk3_0))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,X1)))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_71]), c_0_65]), c_0_48])).
cnf(c_0_74, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_28, c_0_72])).
cnf(c_0_75, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_21]), c_0_31]), c_0_43])]), c_0_74]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 76
# Proof object clause steps            : 47
# Proof object formula steps           : 29
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 24
# Proof object simplifying inferences  : 26
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 31
# Processed clauses                    : 4529
# ...of these trivial                  : 0
# ...subsumed                          : 3435
# ...remaining for further processing  : 1094
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 41
# Backward-rewritten                   : 5
# Generated clauses                    : 13544
# ...of the previous two non-trivial   : 13330
# Contextual simplify-reflections      : 131
# Paramodulations                      : 13537
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1017
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 10
#    Non-unit-clauses                  : 994
# Current number of unprocessed clauses: 8792
# ...number of literals in the above   : 74079
# Current number of archived formulas  : 0
# Current number of archived clauses   : 76
# Clause-clause subsumption calls (NU) : 226590
# Rec. Clause-clause subsumption calls : 34304
# Non-unit clause-clause subsumptions  : 2993
# Unit Clause-clause subsumption calls : 197
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 440428

# -------------------------------------------------
# User time                : 0.381 s
# System time              : 0.003 s
# Total time               : 0.384 s
# Maximum resident set size: 28572 pages
