# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t15_mcart_1, conjecture, ![X1, X2, X3, X4]:(r2_hidden(X1,k2_zfmisc_1(k2_tarski(X2,X3),X4))=>((k1_mcart_1(X1)=X2|k1_mcart_1(X1)=X3)&r2_hidden(k2_mcart_1(X1),X4))), file('data/MPTP/proofs/11039c324a104add880d7243537fa1eb.E_input', t15_mcart_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/11039c324a104add880d7243537fa1eb.E_input', d5_xboole_0)).
fof(t144_zfmisc_1, axiom, ![X1, X2, X3]:~((~(r2_hidden(X1,X3))&(~(r2_hidden(X2,X3))&~(X3=k4_xboole_0(X3,k2_tarski(X1,X2)))))), file('data/MPTP/proofs/11039c324a104add880d7243537fa1eb.E_input', t144_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/11039c324a104add880d7243537fa1eb.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/11039c324a104add880d7243537fa1eb.E_input', t71_enumset1)).
fof(t10_mcart_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(r2_hidden(k1_mcart_1(X1),X2)&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/11039c324a104add880d7243537fa1eb.E_input', t10_mcart_1)).
fof(t38_zfmisc_1, axiom, ![X1, X2, X3]:(r1_tarski(k2_tarski(X1,X2),X3)<=>(r2_hidden(X1,X3)&r2_hidden(X2,X3))), file('data/MPTP/proofs/11039c324a104add880d7243537fa1eb.E_input', t38_zfmisc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/11039c324a104add880d7243537fa1eb.E_input', d10_xboole_0)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/11039c324a104add880d7243537fa1eb.E_input', t7_mcart_1)).
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/11039c324a104add880d7243537fa1eb.E_input', d2_zfmisc_1)).
fof(t12_mcart_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(k1_tarski(X2),X3))=>(k1_mcart_1(X1)=X2&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/11039c324a104add880d7243537fa1eb.E_input', t12_mcart_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/11039c324a104add880d7243537fa1eb.E_input', t69_enumset1)).
fof(c_0_12, negated_conjecture, ~(![X1, X2, X3, X4]:(r2_hidden(X1,k2_zfmisc_1(k2_tarski(X2,X3),X4))=>((k1_mcart_1(X1)=X2|k1_mcart_1(X1)=X3)&r2_hidden(k2_mcart_1(X1),X4)))), inference(assume_negation,[status(cth)],[t15_mcart_1])).
fof(c_0_13, plain, ![X43, X44, X45, X46, X47, X48, X49, X50]:((((r2_hidden(X46,X43)|~r2_hidden(X46,X45)|X45!=k4_xboole_0(X43,X44))&(~r2_hidden(X46,X44)|~r2_hidden(X46,X45)|X45!=k4_xboole_0(X43,X44)))&(~r2_hidden(X47,X43)|r2_hidden(X47,X44)|r2_hidden(X47,X45)|X45!=k4_xboole_0(X43,X44)))&((~r2_hidden(esk10_3(X48,X49,X50),X50)|(~r2_hidden(esk10_3(X48,X49,X50),X48)|r2_hidden(esk10_3(X48,X49,X50),X49))|X50=k4_xboole_0(X48,X49))&((r2_hidden(esk10_3(X48,X49,X50),X48)|r2_hidden(esk10_3(X48,X49,X50),X50)|X50=k4_xboole_0(X48,X49))&(~r2_hidden(esk10_3(X48,X49,X50),X49)|r2_hidden(esk10_3(X48,X49,X50),X50)|X50=k4_xboole_0(X48,X49))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_14, plain, ![X19, X20, X21]:(r2_hidden(X19,X21)|(r2_hidden(X20,X21)|X21=k4_xboole_0(X21,k2_tarski(X19,X20)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t144_zfmisc_1])])])).
fof(c_0_15, plain, ![X54, X55]:k1_enumset1(X54,X54,X55)=k2_tarski(X54,X55), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_16, plain, ![X16, X17, X18]:k2_enumset1(X16,X16,X17,X18)=k1_enumset1(X16,X17,X18), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_17, negated_conjecture, (r2_hidden(esk1_0,k2_zfmisc_1(k2_tarski(esk2_0,esk3_0),esk4_0))&((k1_mcart_1(esk1_0)!=esk2_0|~r2_hidden(k2_mcart_1(esk1_0),esk4_0))&(k1_mcart_1(esk1_0)!=esk3_0|~r2_hidden(k2_mcart_1(esk1_0),esk4_0)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).
cnf(c_0_18, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (r2_hidden(X1,X2)|r2_hidden(X3,X2)|X2=k4_xboole_0(X2,k2_tarski(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, plain, ![X13, X14, X15]:((r2_hidden(k1_mcart_1(X13),X14)|~r2_hidden(X13,k2_zfmisc_1(X14,X15)))&(r2_hidden(k2_mcart_1(X13),X15)|~r2_hidden(X13,k2_zfmisc_1(X14,X15)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_mcart_1])])])).
cnf(c_0_23, negated_conjecture, (r2_hidden(esk1_0,k2_zfmisc_1(k2_tarski(esk2_0,esk3_0),esk4_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (X2=k4_xboole_0(X2,k2_enumset1(X1,X1,X1,X3))|r2_hidden(X3,X2)|r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_20]), c_0_21])).
cnf(c_0_26, plain, (r2_hidden(k1_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27, negated_conjecture, (r2_hidden(esk1_0,k2_zfmisc_1(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0),esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_20]), c_0_21])).
fof(c_0_28, plain, ![X40, X41, X42]:(((r2_hidden(X40,X42)|~r1_tarski(k2_tarski(X40,X41),X42))&(r2_hidden(X41,X42)|~r1_tarski(k2_tarski(X40,X41),X42)))&(~r2_hidden(X40,X42)|~r2_hidden(X41,X42)|r1_tarski(k2_tarski(X40,X41),X42))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t38_zfmisc_1])])])).
fof(c_0_29, plain, ![X52, X53]:(((r1_tarski(X52,X53)|X52!=X53)&(r1_tarski(X53,X52)|X52!=X53))&(~r1_tarski(X52,X53)|~r1_tarski(X53,X52)|X52=X53)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_30, plain, (r2_hidden(X1,X2)|r2_hidden(X3,X2)|~r2_hidden(X4,k2_enumset1(X1,X1,X1,X3))|~r2_hidden(X4,X2)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31, negated_conjecture, (r2_hidden(k1_mcart_1(esk1_0),k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_tarski(X1,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_34, plain, ![X11, X12]:(k1_mcart_1(k4_tarski(X11,X12))=X11&k2_mcart_1(k4_tarski(X11,X12))=X12), inference(variable_rename,[status(thm)],[t7_mcart_1])).
fof(c_0_35, plain, ![X23, X24, X25, X26, X29, X30, X31, X32, X33, X34, X36, X37]:((((r2_hidden(esk5_4(X23,X24,X25,X26),X23)|~r2_hidden(X26,X25)|X25!=k2_zfmisc_1(X23,X24))&((r2_hidden(esk6_4(X23,X24,X25,X26),X24)|~r2_hidden(X26,X25)|X25!=k2_zfmisc_1(X23,X24))&(X26=k4_tarski(esk5_4(X23,X24,X25,X26),esk6_4(X23,X24,X25,X26))|~r2_hidden(X26,X25)|X25!=k2_zfmisc_1(X23,X24))))&(~r2_hidden(X30,X23)|(~r2_hidden(X31,X24)|X29!=k4_tarski(X30,X31))|r2_hidden(X29,X25)|X25!=k2_zfmisc_1(X23,X24)))&((~r2_hidden(esk7_3(X32,X33,X34),X34)|(~r2_hidden(X36,X32)|(~r2_hidden(X37,X33)|esk7_3(X32,X33,X34)!=k4_tarski(X36,X37)))|X34=k2_zfmisc_1(X32,X33))&((r2_hidden(esk8_3(X32,X33,X34),X32)|r2_hidden(esk7_3(X32,X33,X34),X34)|X34=k2_zfmisc_1(X32,X33))&((r2_hidden(esk9_3(X32,X33,X34),X33)|r2_hidden(esk7_3(X32,X33,X34),X34)|X34=k2_zfmisc_1(X32,X33))&(esk7_3(X32,X33,X34)=k4_tarski(esk8_3(X32,X33,X34),esk9_3(X32,X33,X34))|r2_hidden(esk7_3(X32,X33,X34),X34)|X34=k2_zfmisc_1(X32,X33)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
cnf(c_0_36, negated_conjecture, (r2_hidden(esk3_0,X1)|r2_hidden(esk2_0,X1)|~r2_hidden(k1_mcart_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_37, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_enumset1(X1,X1,X1,X3),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_20]), c_0_21])).
cnf(c_0_38, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_33])).
cnf(c_0_39, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_40, plain, (X1=k4_tarski(esk5_4(X2,X3,X4,X1),esk6_4(X2,X3,X4,X1))|~r2_hidden(X1,X4)|X4!=k2_zfmisc_1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41, negated_conjecture, (r2_hidden(esk2_0,k4_xboole_0(X1,X2))|~r2_hidden(k1_mcart_1(esk1_0),k4_xboole_0(X1,X2))|~r2_hidden(esk3_0,X2)), inference(spm,[status(thm)],[c_0_24, c_0_36])).
cnf(c_0_42, plain, (r2_hidden(X1,k2_enumset1(X1,X1,X1,X2))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_43, plain, (esk5_4(X1,X2,X3,X4)=k1_mcart_1(X4)|X3!=k2_zfmisc_1(X1,X2)|~r2_hidden(X4,X3)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_44, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk2_0,k4_xboole_0(X1,k2_enumset1(esk3_0,esk3_0,esk3_0,X2)))|~r2_hidden(k1_mcart_1(esk1_0),k4_xboole_0(X1,k2_enumset1(esk3_0,esk3_0,esk3_0,X2)))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_47, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_48, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_tarski(X3,X1),X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_49, plain, (k4_tarski(k1_mcart_1(X1),esk6_4(X2,X3,X4,X1))=X1|X4!=k2_zfmisc_1(X2,X3)|~r2_hidden(X1,X4)), inference(spm,[status(thm)],[c_0_40, c_0_43])).
cnf(c_0_50, plain, (esk6_4(X1,X2,X3,X4)=k2_mcart_1(X4)|X3!=k2_zfmisc_1(X1,X2)|~r2_hidden(X4,X3)), inference(spm,[status(thm)],[c_0_44, c_0_40])).
cnf(c_0_51, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_45])).
cnf(c_0_52, negated_conjecture, (~r2_hidden(k1_mcart_1(esk1_0),k4_xboole_0(X1,k2_enumset1(esk3_0,esk3_0,esk3_0,X2)))|~r2_hidden(esk2_0,k2_enumset1(esk3_0,esk3_0,esk3_0,X2))), inference(spm,[status(thm)],[c_0_24, c_0_46])).
cnf(c_0_53, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_47])).
cnf(c_0_54, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_enumset1(X3,X3,X3,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_20]), c_0_21])).
cnf(c_0_55, plain, (k4_tarski(k1_mcart_1(X1),k2_mcart_1(X1))=X1|X2!=k2_zfmisc_1(X3,X4)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_56, negated_conjecture, (r2_hidden(esk2_0,X1)|~r2_hidden(k1_mcart_1(esk1_0),k4_xboole_0(X1,k2_enumset1(esk3_0,esk3_0,esk3_0,X2)))), inference(spm,[status(thm)],[c_0_51, c_0_46])).
cnf(c_0_57, negated_conjecture, (r2_hidden(k1_mcart_1(esk1_0),k2_enumset1(esk3_0,esk3_0,esk3_0,X1))|~r2_hidden(esk2_0,k2_enumset1(esk3_0,esk3_0,esk3_0,X1))|~r2_hidden(k1_mcart_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_58, plain, (r2_hidden(X1,k2_enumset1(X2,X2,X2,X1))), inference(spm,[status(thm)],[c_0_54, c_0_38])).
cnf(c_0_59, plain, (r2_hidden(X5,X6)|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)|X5!=k4_tarski(X1,X3)|X6!=k2_zfmisc_1(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_60, plain, (k4_tarski(k1_mcart_1(X1),k2_mcart_1(X1))=X1|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(er,[status(thm)],[c_0_55])).
cnf(c_0_61, negated_conjecture, (r2_hidden(k1_mcart_1(esk1_0),k2_enumset1(esk3_0,esk3_0,esk3_0,X1))|r2_hidden(esk2_0,X2)|~r2_hidden(k1_mcart_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_56, c_0_53])).
cnf(c_0_62, negated_conjecture, (r2_hidden(k1_mcart_1(esk1_0),k2_enumset1(esk3_0,esk3_0,esk3_0,esk2_0))|~r2_hidden(k1_mcart_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
fof(c_0_63, plain, ![X56, X57, X58]:((k1_mcart_1(X56)=X57|~r2_hidden(X56,k2_zfmisc_1(k1_tarski(X57),X58)))&(r2_hidden(k2_mcart_1(X56),X58)|~r2_hidden(X56,k2_zfmisc_1(k1_tarski(X57),X58)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_mcart_1])])])).
fof(c_0_64, plain, ![X22]:k2_tarski(X22,X22)=k1_tarski(X22), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_65, plain, (r2_hidden(X1,k2_zfmisc_1(X2,X3))|X1!=k4_tarski(X4,X5)|~r2_hidden(X5,X3)|~r2_hidden(X4,X2)), inference(er,[status(thm)],[c_0_59])).
cnf(c_0_66, negated_conjecture, (k4_tarski(k1_mcart_1(esk1_0),k2_mcart_1(esk1_0))=esk1_0), inference(spm,[status(thm)],[c_0_60, c_0_27])).
cnf(c_0_67, plain, (r2_hidden(k2_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_68, negated_conjecture, (r2_hidden(k1_mcart_1(esk1_0),k2_enumset1(esk3_0,esk3_0,esk3_0,X1))|r2_hidden(esk2_0,k4_xboole_0(X2,X3))|r2_hidden(k1_mcart_1(esk1_0),X3)|~r2_hidden(k1_mcart_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_61, c_0_53])).
cnf(c_0_69, negated_conjecture, (r2_hidden(k1_mcart_1(esk1_0),k2_enumset1(esk3_0,esk3_0,esk3_0,esk2_0))), inference(spm,[status(thm)],[c_0_62, c_0_58])).
cnf(c_0_70, plain, (k1_mcart_1(X1)=X2|~r2_hidden(X1,k2_zfmisc_1(k1_tarski(X2),X3))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_71, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_72, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(X2,X3))|X1!=esk1_0|~r2_hidden(k2_mcart_1(esk1_0),X3)|~r2_hidden(k1_mcart_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_73, negated_conjecture, (r2_hidden(k2_mcart_1(esk1_0),esk4_0)), inference(spm,[status(thm)],[c_0_67, c_0_27])).
cnf(c_0_74, negated_conjecture, (r2_hidden(esk2_0,k4_xboole_0(k2_enumset1(esk3_0,esk3_0,esk3_0,esk2_0),X1))|r2_hidden(k1_mcart_1(esk1_0),k2_enumset1(esk3_0,esk3_0,esk3_0,X2))|r2_hidden(k1_mcart_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_75, plain, (k1_mcart_1(X1)=X2|~r2_hidden(X1,k2_zfmisc_1(k2_enumset1(X2,X2,X2,X2),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_71]), c_0_20]), c_0_21])).
cnf(c_0_76, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(X2,esk4_0))|X1!=esk1_0|~r2_hidden(k1_mcart_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_77, negated_conjecture, (r2_hidden(k1_mcart_1(esk1_0),k2_enumset1(esk3_0,esk3_0,esk3_0,X1))|r2_hidden(k1_mcart_1(esk1_0),X2)|~r2_hidden(esk2_0,X2)), inference(spm,[status(thm)],[c_0_24, c_0_74])).
cnf(c_0_78, negated_conjecture, (k1_mcart_1(X1)=X2|X1!=esk1_0|~r2_hidden(k1_mcart_1(esk1_0),k2_enumset1(X2,X2,X2,X2))), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_79, negated_conjecture, (r2_hidden(k1_mcart_1(esk1_0),k2_enumset1(esk2_0,esk2_0,esk2_0,X1))|r2_hidden(k1_mcart_1(esk1_0),k2_enumset1(esk3_0,esk3_0,esk3_0,X2))), inference(spm,[status(thm)],[c_0_77, c_0_42])).
cnf(c_0_80, negated_conjecture, (k1_mcart_1(esk1_0)!=esk3_0|~r2_hidden(k2_mcart_1(esk1_0),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_81, negated_conjecture, (k1_mcart_1(X1)=esk3_0|r2_hidden(k1_mcart_1(esk1_0),k2_enumset1(esk2_0,esk2_0,esk2_0,X2))|X1!=esk1_0), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_82, negated_conjecture, (esk3_0!=k1_mcart_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_80, c_0_73])])).
cnf(c_0_83, negated_conjecture, (k1_mcart_1(esk1_0)!=esk2_0|~r2_hidden(k2_mcart_1(esk1_0),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_84, negated_conjecture, (r2_hidden(k1_mcart_1(esk1_0),k2_enumset1(esk2_0,esk2_0,esk2_0,X1))), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_81]), c_0_82])).
cnf(c_0_85, negated_conjecture, (esk2_0!=k1_mcart_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_83, c_0_73])])).
cnf(c_0_86, negated_conjecture, (k1_mcart_1(X1)=esk2_0|X1!=esk1_0), inference(spm,[status(thm)],[c_0_78, c_0_84])).
cnf(c_0_87, negated_conjecture, (k1_mcart_1(X1)!=k1_mcart_1(esk1_0)|X1!=esk1_0), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_88, negated_conjecture, ($false), inference(er,[status(thm)],[c_0_87]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 89
# Proof object clause steps            : 64
# Proof object formula steps           : 25
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 36
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 30
# Processed clauses                    : 3605
# ...of these trivial                  : 55
# ...subsumed                          : 3027
# ...remaining for further processing  : 523
# Other redundant clauses eliminated   : 29
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 34
# Backward-rewritten                   : 21
# Generated clauses                    : 20498
# ...of the previous two non-trivial   : 17394
# Contextual simplify-reflections      : 1
# Paramodulations                      : 20299
# Factorizations                       : 124
# Equation resolutions                 : 75
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 438
#    Positive orientable unit clauses  : 21
#    Positive unorientable unit clauses: 10
#    Negative unit clauses             : 37
#    Non-unit-clauses                  : 370
# Current number of unprocessed clauses: 13695
# ...number of literals in the above   : 63418
# Current number of archived formulas  : 0
# Current number of archived clauses   : 86
# Clause-clause subsumption calls (NU) : 39644
# Rec. Clause-clause subsumption calls : 13629
# Non-unit clause-clause subsumptions  : 1236
# Unit Clause-clause subsumption calls : 1391
# Rewrite failures with RHS unbound    : 943
# BW rewrite match attempts            : 254
# BW rewrite match successes           : 70
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 413395

# -------------------------------------------------
# User time                : 0.246 s
# System time              : 0.010 s
# Total time               : 0.256 s
# Maximum resident set size: 28608 pages
