# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t63_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((~(v2_struct_0(X2))&m1_yellow_0(X2,X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(X5=X4=>((r1_lattice3(X2,X3,X5)=>r1_lattice3(X1,X3,X4))&(r2_lattice3(X2,X3,X5)=>r2_lattice3(X1,X3,X4))))))))), file('data/MPTP/proofs/11259b0d0f65460d9a9e64f0caf1d94c.E_input', t63_yellow_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/11259b0d0f65460d9a9e64f0caf1d94c.E_input', t4_subset)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/11259b0d0f65460d9a9e64f0caf1d94c.E_input', d9_lattice3)).
fof(dt_m1_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_yellow_0(X2,X1)=>l1_orders_2(X2))), file('data/MPTP/proofs/11259b0d0f65460d9a9e64f0caf1d94c.E_input', dt_m1_yellow_0)).
fof(d13_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(m1_yellow_0(X2,X1)<=>(r1_tarski(u1_struct_0(X2),u1_struct_0(X1))&r1_tarski(u1_orders_2(X2),u1_orders_2(X1)))))), file('data/MPTP/proofs/11259b0d0f65460d9a9e64f0caf1d94c.E_input', d13_yellow_0)).
fof(d9_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)<=>r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1)))))), file('data/MPTP/proofs/11259b0d0f65460d9a9e64f0caf1d94c.E_input', d9_orders_2)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/11259b0d0f65460d9a9e64f0caf1d94c.E_input', t3_subset)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/11259b0d0f65460d9a9e64f0caf1d94c.E_input', l3_subset_1)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/11259b0d0f65460d9a9e64f0caf1d94c.E_input', d8_lattice3)).
fof(c_0_9, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((~(v2_struct_0(X2))&m1_yellow_0(X2,X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(X5=X4=>((r1_lattice3(X2,X3,X5)=>r1_lattice3(X1,X3,X4))&(r2_lattice3(X2,X3,X5)=>r2_lattice3(X1,X3,X4)))))))))), inference(assume_negation,[status(cth)],[t63_yellow_0])).
fof(c_0_10, plain, ![X11, X12, X13]:(~r2_hidden(X11,X12)|~m1_subset_1(X12,k1_zfmisc_1(X13))|m1_subset_1(X11,X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_11, negated_conjecture, ((~v2_struct_0(esk1_0)&l1_orders_2(esk1_0))&((~v2_struct_0(esk2_0)&m1_yellow_0(esk2_0,esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&(m1_subset_1(esk4_0,u1_struct_0(esk1_0))&(m1_subset_1(esk5_0,u1_struct_0(esk2_0))&(esk5_0=esk4_0&(((r2_lattice3(esk2_0,esk3_0,esk5_0)|r1_lattice3(esk2_0,esk3_0,esk5_0))&(~r2_lattice3(esk1_0,esk3_0,esk4_0)|r1_lattice3(esk2_0,esk3_0,esk5_0)))&((r2_lattice3(esk2_0,esk3_0,esk5_0)|~r1_lattice3(esk1_0,esk3_0,esk4_0))&(~r2_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk1_0,esk3_0,esk4_0)))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])])).
fof(c_0_12, plain, ![X14, X15, X16, X17]:((~r2_lattice3(X14,X15,X16)|(~m1_subset_1(X17,u1_struct_0(X14))|(~r2_hidden(X17,X15)|r1_orders_2(X14,X17,X16)))|~m1_subset_1(X16,u1_struct_0(X14))|~l1_orders_2(X14))&((m1_subset_1(esk6_3(X14,X15,X16),u1_struct_0(X14))|r2_lattice3(X14,X15,X16)|~m1_subset_1(X16,u1_struct_0(X14))|~l1_orders_2(X14))&((r2_hidden(esk6_3(X14,X15,X16),X15)|r2_lattice3(X14,X15,X16)|~m1_subset_1(X16,u1_struct_0(X14))|~l1_orders_2(X14))&(~r1_orders_2(X14,esk6_3(X14,X15,X16),X16)|r2_lattice3(X14,X15,X16)|~m1_subset_1(X16,u1_struct_0(X14))|~l1_orders_2(X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
cnf(c_0_13, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_18, plain, ![X19, X20]:(~l1_orders_2(X19)|(~m1_yellow_0(X20,X19)|l1_orders_2(X20))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_yellow_0])])])).
cnf(c_0_19, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_20, negated_conjecture, (r2_hidden(esk6_3(esk1_0,X1,esk4_0),X1)|r2_lattice3(esk1_0,X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])])).
cnf(c_0_21, plain, (l1_orders_2(X2)|~l1_orders_2(X1)|~m1_yellow_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, negated_conjecture, (m1_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_23, plain, ![X21, X22]:(((r1_tarski(u1_struct_0(X22),u1_struct_0(X21))|~m1_yellow_0(X22,X21)|~l1_orders_2(X22)|~l1_orders_2(X21))&(r1_tarski(u1_orders_2(X22),u1_orders_2(X21))|~m1_yellow_0(X22,X21)|~l1_orders_2(X22)|~l1_orders_2(X21)))&(~r1_tarski(u1_struct_0(X22),u1_struct_0(X21))|~r1_tarski(u1_orders_2(X22),u1_orders_2(X21))|m1_yellow_0(X22,X21)|~l1_orders_2(X22)|~l1_orders_2(X21))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d13_yellow_0])])])])).
cnf(c_0_24, plain, (r1_orders_2(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_25, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,esk4_0)|m1_subset_1(esk6_3(esk1_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_26, negated_conjecture, (l1_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_17])])).
cnf(c_0_27, plain, (r1_tarski(u1_orders_2(X1),u1_orders_2(X2))|~m1_yellow_0(X1,X2)|~l1_orders_2(X1)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_28, plain, ![X23, X24, X25]:((~r1_orders_2(X23,X24,X25)|r2_hidden(k4_tarski(X24,X25),u1_orders_2(X23))|~m1_subset_1(X25,u1_struct_0(X23))|~m1_subset_1(X24,u1_struct_0(X23))|~l1_orders_2(X23))&(~r2_hidden(k4_tarski(X24,X25),u1_orders_2(X23))|r1_orders_2(X23,X24,X25)|~m1_subset_1(X25,u1_struct_0(X23))|~m1_subset_1(X24,u1_struct_0(X23))|~l1_orders_2(X23))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_orders_2])])])])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_30, negated_conjecture, (esk5_0=esk4_0), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_31, negated_conjecture, (r1_orders_2(esk2_0,esk6_3(esk1_0,esk3_0,esk4_0),X1)|r2_lattice3(esk1_0,esk3_0,esk4_0)|~r2_hidden(esk6_3(esk1_0,esk3_0,esk4_0),X2)|~r2_lattice3(esk2_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26])])).
cnf(c_0_32, negated_conjecture, (r2_lattice3(esk2_0,esk3_0,esk5_0)|r1_lattice3(esk2_0,esk3_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_33, negated_conjecture, (r1_lattice3(esk2_0,esk3_0,esk5_0)|~r2_lattice3(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_34, plain, ![X26, X27]:((~m1_subset_1(X26,k1_zfmisc_1(X27))|r1_tarski(X26,X27))&(~r1_tarski(X26,X27)|m1_subset_1(X26,k1_zfmisc_1(X27)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_35, plain, (r1_tarski(u1_orders_2(X1),u1_orders_2(X2))|~m1_yellow_0(X1,X2)|~l1_orders_2(X2)), inference(csr,[status(thm)],[c_0_27, c_0_21])).
cnf(c_0_36, plain, (r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(rw,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_38, negated_conjecture, (r1_orders_2(esk2_0,esk6_3(esk1_0,esk3_0,esk4_0),X1)|r2_lattice3(esk1_0,esk3_0,esk4_0)|~r2_lattice3(esk2_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_31, c_0_20])).
cnf(c_0_39, negated_conjecture, (r2_lattice3(esk2_0,esk3_0,esk4_0)|r1_lattice3(esk2_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_30]), c_0_30])).
cnf(c_0_40, negated_conjecture, (r1_lattice3(esk2_0,esk3_0,esk4_0)|~r2_lattice3(esk1_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_33, c_0_30])).
fof(c_0_41, plain, ![X28, X29, X30]:(~m1_subset_1(X29,k1_zfmisc_1(X28))|(~r2_hidden(X30,X29)|r2_hidden(X30,X28))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
cnf(c_0_42, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43, negated_conjecture, (r1_tarski(u1_orders_2(esk2_0),u1_orders_2(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_22]), c_0_17])])).
cnf(c_0_44, negated_conjecture, (r2_hidden(k4_tarski(X1,esk4_0),u1_orders_2(esk2_0))|~r1_orders_2(esk2_0,X1,esk4_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_26])])).
cnf(c_0_45, negated_conjecture, (r1_orders_2(esk2_0,esk6_3(esk1_0,esk3_0,esk4_0),esk4_0)|r1_lattice3(esk2_0,esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_37])]), c_0_40])).
fof(c_0_46, plain, ![X31, X32, X33, X34]:((~r1_lattice3(X31,X32,X33)|(~m1_subset_1(X34,u1_struct_0(X31))|(~r2_hidden(X34,X32)|r1_orders_2(X31,X33,X34)))|~m1_subset_1(X33,u1_struct_0(X31))|~l1_orders_2(X31))&((m1_subset_1(esk7_3(X31,X32,X33),u1_struct_0(X31))|r1_lattice3(X31,X32,X33)|~m1_subset_1(X33,u1_struct_0(X31))|~l1_orders_2(X31))&((r2_hidden(esk7_3(X31,X32,X33),X32)|r1_lattice3(X31,X32,X33)|~m1_subset_1(X33,u1_struct_0(X31))|~l1_orders_2(X31))&(~r1_orders_2(X31,X33,esk7_3(X31,X32,X33))|r1_lattice3(X31,X32,X33)|~m1_subset_1(X33,u1_struct_0(X31))|~l1_orders_2(X31))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
cnf(c_0_47, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_48, negated_conjecture, (m1_subset_1(u1_orders_2(esk2_0),k1_zfmisc_1(u1_orders_2(esk1_0)))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_49, negated_conjecture, (r2_hidden(k4_tarski(esk6_3(esk1_0,esk3_0,esk4_0),esk4_0),u1_orders_2(esk2_0))|r1_lattice3(esk2_0,esk3_0,esk4_0)|~m1_subset_1(esk6_3(esk1_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_50, plain, (r2_hidden(esk7_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_51, negated_conjecture, (r2_hidden(X1,u1_orders_2(esk1_0))|~r2_hidden(X1,u1_orders_2(esk2_0))), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_52, negated_conjecture, (r2_hidden(k4_tarski(esk6_3(esk1_0,esk3_0,esk4_0),esk4_0),u1_orders_2(esk2_0))|r1_lattice3(esk2_0,esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_25]), c_0_40])).
cnf(c_0_53, negated_conjecture, (r2_hidden(esk7_3(esk1_0,X1,esk4_0),X1)|r1_lattice3(esk1_0,X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_16]), c_0_17])])).
cnf(c_0_54, plain, (r1_orders_2(X3,X1,X2)|~r2_hidden(k4_tarski(X1,X2),u1_orders_2(X3))|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_55, negated_conjecture, (r2_hidden(k4_tarski(esk6_3(esk1_0,esk3_0,esk4_0),esk4_0),u1_orders_2(esk1_0))|r1_lattice3(esk2_0,esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_56, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_57, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_58, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk4_0)|m1_subset_1(esk7_3(esk1_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_19, c_0_53])).
cnf(c_0_59, negated_conjecture, (r1_orders_2(esk1_0,esk6_3(esk1_0,esk3_0,esk4_0),esk4_0)|r1_lattice3(esk2_0,esk3_0,esk4_0)|~m1_subset_1(esk6_3(esk1_0,esk3_0,esk4_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_16]), c_0_17])])).
cnf(c_0_60, negated_conjecture, (r2_lattice3(esk1_0,X1,esk4_0)|m1_subset_1(esk6_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_16]), c_0_17])])).
cnf(c_0_61, negated_conjecture, (r1_orders_2(esk2_0,X1,esk7_3(esk1_0,esk3_0,esk4_0))|r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_hidden(esk7_3(esk1_0,esk3_0,esk4_0),X2)|~r1_lattice3(esk2_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_26])])).
cnf(c_0_62, plain, (r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk6_3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_63, negated_conjecture, (r1_orders_2(esk1_0,esk6_3(esk1_0,esk3_0,esk4_0),esk4_0)|r1_lattice3(esk2_0,esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_40])).
cnf(c_0_64, negated_conjecture, (r1_orders_2(esk2_0,X1,esk7_3(esk1_0,esk3_0,esk4_0))|r1_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk2_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_61, c_0_53])).
cnf(c_0_65, negated_conjecture, (r1_lattice3(esk2_0,esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_16]), c_0_17])]), c_0_40])).
cnf(c_0_66, negated_conjecture, (r2_hidden(k4_tarski(X1,esk7_3(esk1_0,esk3_0,esk4_0)),u1_orders_2(esk2_0))|r1_lattice3(esk1_0,esk3_0,esk4_0)|~r1_orders_2(esk2_0,X1,esk7_3(esk1_0,esk3_0,esk4_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_58]), c_0_26])])).
cnf(c_0_67, negated_conjecture, (r1_orders_2(esk2_0,esk4_0,esk7_3(esk1_0,esk3_0,esk4_0))|r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_37])])).
cnf(c_0_68, negated_conjecture, (r2_hidden(k4_tarski(esk4_0,esk7_3(esk1_0,esk3_0,esk4_0)),u1_orders_2(esk2_0))|r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_37])])).
cnf(c_0_69, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_70, negated_conjecture, (r2_hidden(k4_tarski(esk4_0,esk7_3(esk1_0,esk3_0,esk4_0)),u1_orders_2(esk1_0))|r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_51, c_0_68])).
cnf(c_0_71, negated_conjecture, (r1_lattice3(esk1_0,X1,esk4_0)|m1_subset_1(esk7_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_16]), c_0_17])])).
cnf(c_0_72, negated_conjecture, (r2_lattice3(esk2_0,esk3_0,esk5_0)|~r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_73, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk7_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_74, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,esk7_3(esk1_0,esk3_0,esk4_0))|r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_70]), c_0_16]), c_0_17])]), c_0_71])).
cnf(c_0_75, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_yellow_0(X1,X2)|~l1_orders_2(X1)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_76, negated_conjecture, (r2_lattice3(esk2_0,esk3_0,esk4_0)|~r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_72, c_0_30])).
cnf(c_0_77, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_16]), c_0_17])])).
cnf(c_0_78, negated_conjecture, (~r2_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_79, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_yellow_0(X1,X2)|~l1_orders_2(X2)), inference(csr,[status(thm)],[c_0_75, c_0_21])).
cnf(c_0_80, negated_conjecture, (r2_lattice3(esk2_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_77])])).
cnf(c_0_81, negated_conjecture, (~r2_lattice3(esk1_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_77])])).
cnf(c_0_82, negated_conjecture, (r1_tarski(u1_struct_0(esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_22]), c_0_17])])).
cnf(c_0_83, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_47, c_0_14])).
cnf(c_0_84, negated_conjecture, (r1_orders_2(esk2_0,esk6_3(esk1_0,esk3_0,esk4_0),esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_80]), c_0_37])]), c_0_81])).
cnf(c_0_85, negated_conjecture, (m1_subset_1(esk6_3(esk1_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[c_0_25, c_0_81])).
cnf(c_0_86, negated_conjecture, (m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_42, c_0_82])).
cnf(c_0_87, negated_conjecture, (r2_hidden(esk6_3(esk1_0,esk3_0,esk4_0),u1_struct_0(esk2_0))|r2_lattice3(esk1_0,esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_83, c_0_20])).
cnf(c_0_88, negated_conjecture, (r2_hidden(k4_tarski(esk6_3(esk1_0,esk3_0,esk4_0),esk4_0),u1_orders_2(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_84]), c_0_85])])).
cnf(c_0_89, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_13, c_0_86])).
cnf(c_0_90, negated_conjecture, (r2_hidden(esk6_3(esk1_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[c_0_87, c_0_81])).
cnf(c_0_91, negated_conjecture, (r2_hidden(k4_tarski(esk6_3(esk1_0,esk3_0,esk4_0),esk4_0),u1_orders_2(esk1_0))), inference(spm,[status(thm)],[c_0_51, c_0_88])).
cnf(c_0_92, negated_conjecture, (m1_subset_1(esk6_3(esk1_0,esk3_0,esk4_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_93, negated_conjecture, (r1_orders_2(esk1_0,esk6_3(esk1_0,esk3_0,esk4_0),esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_91]), c_0_16]), c_0_92]), c_0_17])])).
cnf(c_0_94, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_93]), c_0_16]), c_0_17])]), c_0_81]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 95
# Proof object clause steps            : 76
# Proof object formula steps           : 19
# Proof object conjectures             : 61
# Proof object clause conjectures      : 58
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 40
# Proof object simplifying inferences  : 71
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 354
# ...of these trivial                  : 1
# ...subsumed                          : 23
# ...remaining for further processing  : 330
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 15
# Backward-rewritten                   : 105
# Generated clauses                    : 734
# ...of the previous two non-trivial   : 725
# Contextual simplify-reflections      : 12
# Paramodulations                      : 728
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 174
#    Positive orientable unit clauses  : 26
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 145
# Current number of unprocessed clauses: 413
# ...number of literals in the above   : 1807
# Current number of archived formulas  : 0
# Current number of archived clauses   : 156
# Clause-clause subsumption calls (NU) : 10387
# Rec. Clause-clause subsumption calls : 3936
# Non-unit clause-clause subsumptions  : 50
# Unit Clause-clause subsumption calls : 224
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 24
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 23130

# -------------------------------------------------
# User time                : 0.018 s
# System time              : 0.006 s
# Total time               : 0.024 s
# Maximum resident set size: 28788 pages
