# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t34_funct_2, conjecture, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X2,X1)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))))=>((k2_relset_1(X1,X2,X3)=X2&(v2_funct_1(X3)&![X5, X6]:(((r2_hidden(X5,X2)&k1_funct_1(X4,X5)=X6)=>(r2_hidden(X6,X1)&k1_funct_1(X3,X6)=X5))&((r2_hidden(X6,X1)&k1_funct_1(X3,X6)=X5)=>(r2_hidden(X5,X2)&k1_funct_1(X4,X5)=X6)))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|X4=k2_funct_1(X3)))))), file('data/MPTP/proofs/11590de0c9dd4f2abac61e824b65d3f3.E_input', t34_funct_2)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(((X2=k1_xboole_0=>X1=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X2,X3)))&(X2=k1_xboole_0=>(X1=k1_xboole_0|(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0))))), file('data/MPTP/proofs/11590de0c9dd4f2abac61e824b65d3f3.E_input', d1_funct_2)).
fof(t3_funct_2, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_funct_1(X1)&(v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))))), file('data/MPTP/proofs/11590de0c9dd4f2abac61e824b65d3f3.E_input', t3_funct_2)).
fof(t55_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)=>(k2_relat_1(X1)=k1_relat_1(k2_funct_1(X1))&k1_relat_1(X1)=k2_relat_1(k2_funct_1(X1))))), file('data/MPTP/proofs/11590de0c9dd4f2abac61e824b65d3f3.E_input', t55_funct_1)).
fof(dt_k2_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_relat_1(k2_funct_1(X1))&v1_funct_1(k2_funct_1(X1)))), file('data/MPTP/proofs/11590de0c9dd4f2abac61e824b65d3f3.E_input', dt_k2_funct_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/11590de0c9dd4f2abac61e824b65d3f3.E_input', redefinition_k1_relset_1)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/11590de0c9dd4f2abac61e824b65d3f3.E_input', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/11590de0c9dd4f2abac61e824b65d3f3.E_input', fc6_relat_1)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/11590de0c9dd4f2abac61e824b65d3f3.E_input', redefinition_k2_relset_1)).
fof(t12_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r2_hidden(X1,k1_relat_1(X2))=>r2_hidden(k1_funct_1(X2,X1),k2_relat_1(X2)))), file('data/MPTP/proofs/11590de0c9dd4f2abac61e824b65d3f3.E_input', t12_funct_1)).
fof(t5_funct_2, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X3)=X1&![X4]:(r2_hidden(X4,X1)=>r2_hidden(k1_funct_1(X3,X4),X2)))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))), file('data/MPTP/proofs/11590de0c9dd4f2abac61e824b65d3f3.E_input', t5_funct_2)).
fof(t31_funct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>((v2_funct_1(X3)&k2_relset_1(X1,X2,X3)=X2)=>(v1_funct_1(k2_funct_1(X3))&(v1_funct_2(k2_funct_1(X3),X2,X1)&m1_subset_1(k2_funct_1(X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1))))))), file('data/MPTP/proofs/11590de0c9dd4f2abac61e824b65d3f3.E_input', t31_funct_2)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/11590de0c9dd4f2abac61e824b65d3f3.E_input', d10_xboole_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/11590de0c9dd4f2abac61e824b65d3f3.E_input', t3_subset)).
fof(dt_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k2_relset_1(X1,X2,X3),k1_zfmisc_1(X2))), file('data/MPTP/proofs/11590de0c9dd4f2abac61e824b65d3f3.E_input', dt_k2_relset_1)).
fof(t60_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((v2_funct_1(X1)&(k1_relat_1(X1)=k2_relat_1(X2)&(k2_relat_1(X1)=k1_relat_1(X2)&![X3, X4]:((r2_hidden(X3,k1_relat_1(X1))&r2_hidden(X4,k1_relat_1(X2)))=>(k1_funct_1(X1,X3)=X4<=>k1_funct_1(X2,X4)=X3)))))=>X2=k2_funct_1(X1)))), file('data/MPTP/proofs/11590de0c9dd4f2abac61e824b65d3f3.E_input', t60_funct_1)).
fof(c_0_16, negated_conjecture, ~(![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X2,X1)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))))=>((k2_relset_1(X1,X2,X3)=X2&(v2_funct_1(X3)&![X5, X6]:(((r2_hidden(X5,X2)&k1_funct_1(X4,X5)=X6)=>(r2_hidden(X6,X1)&k1_funct_1(X3,X6)=X5))&((r2_hidden(X6,X1)&k1_funct_1(X3,X6)=X5)=>(r2_hidden(X5,X2)&k1_funct_1(X4,X5)=X6)))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|X4=k2_funct_1(X3))))))), inference(assume_negation,[status(cth)],[t34_funct_2])).
fof(c_0_17, plain, ![X13, X14, X15]:((((~v1_funct_2(X15,X13,X14)|X13=k1_relset_1(X13,X14,X15)|X14=k1_xboole_0|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))&(X13!=k1_relset_1(X13,X14,X15)|v1_funct_2(X15,X13,X14)|X14=k1_xboole_0|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14)))))&((~v1_funct_2(X15,X13,X14)|X13=k1_relset_1(X13,X14,X15)|X13!=k1_xboole_0|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))&(X13!=k1_relset_1(X13,X14,X15)|v1_funct_2(X15,X13,X14)|X13!=k1_xboole_0|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))))&((~v1_funct_2(X15,X13,X14)|X15=k1_xboole_0|X13=k1_xboole_0|X14!=k1_xboole_0|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))&(X15!=k1_xboole_0|v1_funct_2(X15,X13,X14)|X13=k1_xboole_0|X14!=k1_xboole_0|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
fof(c_0_18, negated_conjecture, ![X11, X12]:((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,esk1_0,esk2_0)&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))&((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,esk2_0,esk1_0)&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))))&((k2_relset_1(esk1_0,esk2_0,esk3_0)=esk2_0&(v2_funct_1(esk3_0)&(((r2_hidden(X12,esk1_0)|(~r2_hidden(X11,esk2_0)|k1_funct_1(esk4_0,X11)!=X12))&(k1_funct_1(esk3_0,X12)=X11|(~r2_hidden(X11,esk2_0)|k1_funct_1(esk4_0,X11)!=X12)))&((r2_hidden(X11,esk2_0)|(~r2_hidden(X12,esk1_0)|k1_funct_1(esk3_0,X12)!=X11))&(k1_funct_1(esk4_0,X11)=X12|(~r2_hidden(X12,esk1_0)|k1_funct_1(esk3_0,X12)!=X11))))))&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&esk4_0!=k2_funct_1(esk3_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])])).
fof(c_0_19, plain, ![X16]:((v1_funct_1(X16)|(~v1_relat_1(X16)|~v1_funct_1(X16)))&((v1_funct_2(X16,k1_relat_1(X16),k2_relat_1(X16))|(~v1_relat_1(X16)|~v1_funct_1(X16)))&(m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X16),k2_relat_1(X16))))|(~v1_relat_1(X16)|~v1_funct_1(X16))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_funct_2])])])).
fof(c_0_20, plain, ![X17]:((k2_relat_1(X17)=k1_relat_1(k2_funct_1(X17))|~v2_funct_1(X17)|(~v1_relat_1(X17)|~v1_funct_1(X17)))&(k1_relat_1(X17)=k2_relat_1(k2_funct_1(X17))|~v2_funct_1(X17)|(~v1_relat_1(X17)|~v1_funct_1(X17)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_funct_1])])])).
fof(c_0_21, plain, ![X18]:((v1_relat_1(k2_funct_1(X18))|(~v1_relat_1(X18)|~v1_funct_1(X18)))&(v1_funct_1(k2_funct_1(X18))|(~v1_relat_1(X18)|~v1_funct_1(X18)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_funct_1])])])).
fof(c_0_22, plain, ![X19, X20, X21]:(~m1_subset_1(X21,k1_zfmisc_1(k2_zfmisc_1(X19,X20)))|k1_relset_1(X19,X20,X21)=k1_relat_1(X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_23, plain, (X2=k1_relset_1(X2,X3,X1)|X3=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, negated_conjecture, (v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_27, plain, ![X22, X23]:(~v1_relat_1(X22)|(~m1_subset_1(X23,k1_zfmisc_1(X22))|v1_relat_1(X23))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_28, plain, ![X24, X25]:v1_relat_1(k2_zfmisc_1(X24,X25)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_30, negated_conjecture, (v1_funct_2(esk4_0,esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_31, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_32, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_33, plain, (k1_relat_1(X1)=k2_relat_1(k2_funct_1(X1))|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_34, plain, (v1_funct_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_35, plain, (v1_relat_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_36, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_37, negated_conjecture, (k1_relset_1(esk1_0,esk2_0,esk3_0)=esk1_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_38, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_39, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_40, plain, ![X26, X27, X28]:(~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(X26,X27)))|k2_relset_1(X26,X27,X28)=k2_relat_1(X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
cnf(c_0_41, plain, (v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_42, plain, ![X29, X30]:(~v1_relat_1(X30)|~v1_funct_1(X30)|(~r2_hidden(X29,k1_relat_1(X30))|r2_hidden(k1_funct_1(X30,X29),k2_relat_1(X30)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_funct_1])])).
cnf(c_0_43, negated_conjecture, (k1_funct_1(esk4_0,X1)=X2|~r2_hidden(X2,esk1_0)|k1_funct_1(esk3_0,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_44, negated_conjecture, (k1_relset_1(esk2_0,esk1_0,esk4_0)=esk2_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_45, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X2,esk1_0)|k1_funct_1(esk3_0,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_46, plain, (m1_subset_1(k2_funct_1(X1),k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(k2_funct_1(X1)),k1_relat_1(X1))))|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34]), c_0_35])).
cnf(c_0_47, negated_conjecture, (k1_relat_1(esk3_0)=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_24])])).
cnf(c_0_48, negated_conjecture, (v1_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_24]), c_0_39])])).
cnf(c_0_49, negated_conjecture, (v2_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_50, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_51, negated_conjecture, (k2_relset_1(esk1_0,esk2_0,esk3_0)=esk2_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_52, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_53, plain, (v1_funct_2(k2_funct_1(X1),k1_relat_1(k2_funct_1(X1)),k1_relat_1(X1))|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_33]), c_0_34]), c_0_35])).
fof(c_0_54, plain, ![X31, X32, X33]:(((v1_funct_1(X33)|(r2_hidden(esk5_3(X31,X32,X33),X31)|k1_relat_1(X33)!=X31)|(~v1_relat_1(X33)|~v1_funct_1(X33)))&((v1_funct_2(X33,X31,X32)|(r2_hidden(esk5_3(X31,X32,X33),X31)|k1_relat_1(X33)!=X31)|(~v1_relat_1(X33)|~v1_funct_1(X33)))&(m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X32)))|(r2_hidden(esk5_3(X31,X32,X33),X31)|k1_relat_1(X33)!=X31)|(~v1_relat_1(X33)|~v1_funct_1(X33)))))&((v1_funct_1(X33)|(~r2_hidden(k1_funct_1(X33,esk5_3(X31,X32,X33)),X32)|k1_relat_1(X33)!=X31)|(~v1_relat_1(X33)|~v1_funct_1(X33)))&((v1_funct_2(X33,X31,X32)|(~r2_hidden(k1_funct_1(X33,esk5_3(X31,X32,X33)),X32)|k1_relat_1(X33)!=X31)|(~v1_relat_1(X33)|~v1_funct_1(X33)))&(m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X32)))|(~r2_hidden(k1_funct_1(X33,esk5_3(X31,X32,X33)),X32)|k1_relat_1(X33)!=X31)|(~v1_relat_1(X33)|~v1_funct_1(X33)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_funct_2])])])])).
cnf(c_0_55, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_56, negated_conjecture, (k1_funct_1(esk4_0,k1_funct_1(esk3_0,X1))=X1|~r2_hidden(X1,esk1_0)), inference(er,[status(thm)],[c_0_43])).
cnf(c_0_57, negated_conjecture, (v1_relat_1(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_29]), c_0_39])])).
cnf(c_0_58, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_59, negated_conjecture, (k1_relat_1(esk4_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_44]), c_0_29])])).
cnf(c_0_60, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,X1),esk2_0)|~r2_hidden(X1,esk1_0)), inference(er,[status(thm)],[c_0_45])).
fof(c_0_61, plain, ![X35, X36, X37]:((v1_funct_1(k2_funct_1(X37))|(~v2_funct_1(X37)|k2_relset_1(X35,X36,X37)!=X36)|(~v1_funct_1(X37)|(~v1_funct_2(X37,X35,X36)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(X35,X36))))))&((v1_funct_2(k2_funct_1(X37),X36,X35)|(~v2_funct_1(X37)|k2_relset_1(X35,X36,X37)!=X36)|(~v1_funct_1(X37)|(~v1_funct_2(X37,X35,X36)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(X35,X36))))))&(m1_subset_1(k2_funct_1(X37),k1_zfmisc_1(k2_zfmisc_1(X36,X35)))|(~v2_funct_1(X37)|k2_relset_1(X35,X36,X37)!=X36)|(~v1_funct_1(X37)|(~v1_funct_2(X37,X35,X36)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(X35,X36)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_funct_2])])])).
cnf(c_0_62, negated_conjecture, (m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(k2_funct_1(esk3_0)),esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48]), c_0_49]), c_0_50])])).
cnf(c_0_63, plain, (k2_relat_1(X1)=k1_relat_1(k2_funct_1(X1))|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_64, negated_conjecture, (k2_relat_1(esk3_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_24])])).
cnf(c_0_65, negated_conjecture, (v1_funct_2(k2_funct_1(esk3_0),k1_relat_1(k2_funct_1(esk3_0)),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_47]), c_0_48]), c_0_49]), c_0_50])])).
cnf(c_0_66, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r2_hidden(k1_funct_1(X1,esk5_3(X2,X3,X1)),X3)|k1_relat_1(X1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_67, negated_conjecture, (r2_hidden(X1,k2_relat_1(esk4_0))|~r2_hidden(X1,esk1_0)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57]), c_0_58])]), c_0_59]), c_0_60])).
cnf(c_0_68, plain, (r2_hidden(k1_funct_1(k2_funct_1(X1),X2),k1_relat_1(X1))|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(k2_funct_1(X1)))|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_33]), c_0_34]), c_0_35])).
cnf(c_0_69, plain, (v1_funct_1(k2_funct_1(X1))|~v2_funct_1(X1)|k2_relset_1(X2,X3,X1)!=X3|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_70, negated_conjecture, (m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64]), c_0_48]), c_0_49]), c_0_50])])).
cnf(c_0_71, negated_conjecture, (v1_funct_2(k2_funct_1(esk3_0),esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_63]), c_0_64]), c_0_48]), c_0_49]), c_0_50])])).
cnf(c_0_72, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k2_relat_1(esk4_0))))|k1_relat_1(X1)!=X2|~v1_relat_1(X1)|~r2_hidden(k1_funct_1(X1,esk5_3(X2,k2_relat_1(esk4_0),X1)),esk1_0)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_73, negated_conjecture, (r2_hidden(k1_funct_1(k2_funct_1(esk3_0),X1),esk1_0)|~r2_hidden(X1,k1_relat_1(k2_funct_1(esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_47]), c_0_48]), c_0_49]), c_0_50])])).
cnf(c_0_74, negated_conjecture, (v1_funct_1(k2_funct_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_51]), c_0_49]), c_0_24]), c_0_25]), c_0_50])])).
cnf(c_0_75, negated_conjecture, (k1_relset_1(esk2_0,esk1_0,k2_funct_1(esk3_0))=esk2_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_70]), c_0_71])]), c_0_31])).
cnf(c_0_76, negated_conjecture, (m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(X1,k2_relat_1(esk4_0))))|k1_relat_1(k2_funct_1(esk3_0))!=X1|~v1_relat_1(k2_funct_1(esk3_0))|~r2_hidden(esk5_3(X1,k2_relat_1(esk4_0),k2_funct_1(esk3_0)),k1_relat_1(k2_funct_1(esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_74])])).
cnf(c_0_77, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|r2_hidden(esk5_3(X2,X3,X1),X2)|k1_relat_1(X1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_78, negated_conjecture, (k1_relat_1(k2_funct_1(esk3_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_75]), c_0_70])])).
cnf(c_0_79, negated_conjecture, (v1_relat_1(k2_funct_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_62]), c_0_39])])).
fof(c_0_80, plain, ![X38, X39]:(((r1_tarski(X38,X39)|X38!=X39)&(r1_tarski(X39,X38)|X38!=X39))&(~r1_tarski(X38,X39)|~r1_tarski(X39,X38)|X38=X39)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_81, plain, ![X40, X41]:((~m1_subset_1(X40,k1_zfmisc_1(X41))|r1_tarski(X40,X41))&(~r1_tarski(X40,X41)|m1_subset_1(X40,k1_zfmisc_1(X41)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_82, plain, ![X42, X43, X44]:(~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(X42,X43)))|m1_subset_1(k2_relset_1(X42,X43,X44),k1_zfmisc_1(X43))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_relset_1])])).
cnf(c_0_83, negated_conjecture, (m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(X1,k2_relat_1(esk4_0))))|esk2_0!=X1|~v1_relat_1(k2_funct_1(esk3_0))|~r2_hidden(esk5_3(X1,k2_relat_1(esk4_0),k2_funct_1(esk3_0)),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_63]), c_0_64]), c_0_64]), c_0_48]), c_0_49]), c_0_50])])).
cnf(c_0_84, negated_conjecture, (r2_hidden(esk5_3(X1,X2,k2_funct_1(esk3_0)),X1)|m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|esk2_0!=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_79]), c_0_74])])).
cnf(c_0_85, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_86, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_87, plain, (m1_subset_1(k2_relset_1(X2,X3,X1),k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_88, negated_conjecture, (m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(X1,k2_relat_1(esk4_0))))|esk2_0!=X1|~r2_hidden(esk5_3(X1,k2_relat_1(esk4_0),k2_funct_1(esk3_0)),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_83, c_0_79])])).
cnf(c_0_89, negated_conjecture, (r2_hidden(esk5_3(esk2_0,X1,k2_funct_1(esk3_0)),esk2_0)|m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,X1)))), inference(er,[status(thm)],[c_0_84])).
cnf(c_0_90, plain, (X1=X2|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_91, plain, (m1_subset_1(k2_relat_1(X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(spm,[status(thm)],[c_0_87, c_0_52])).
cnf(c_0_92, negated_conjecture, (m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,k2_relat_1(esk4_0))))), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_93, plain, (X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_90, c_0_86])).
cnf(c_0_94, negated_conjecture, (m1_subset_1(k2_relat_1(esk4_0),k1_zfmisc_1(esk1_0))), inference(spm,[status(thm)],[c_0_91, c_0_29])).
cnf(c_0_95, negated_conjecture, (m1_subset_1(k2_relat_1(k2_funct_1(esk3_0)),k1_zfmisc_1(k2_relat_1(esk4_0)))), inference(spm,[status(thm)],[c_0_91, c_0_92])).
fof(c_0_96, plain, ![X45, X46]:(((r2_hidden(esk6_2(X45,X46),k1_relat_1(X45))|k2_relat_1(X45)!=k1_relat_1(X46)|k1_relat_1(X45)!=k2_relat_1(X46)|~v2_funct_1(X45)|X46=k2_funct_1(X45)|(~v1_relat_1(X46)|~v1_funct_1(X46))|(~v1_relat_1(X45)|~v1_funct_1(X45)))&(r2_hidden(esk7_2(X45,X46),k1_relat_1(X46))|k2_relat_1(X45)!=k1_relat_1(X46)|k1_relat_1(X45)!=k2_relat_1(X46)|~v2_funct_1(X45)|X46=k2_funct_1(X45)|(~v1_relat_1(X46)|~v1_funct_1(X46))|(~v1_relat_1(X45)|~v1_funct_1(X45))))&((k1_funct_1(X45,esk6_2(X45,X46))!=esk7_2(X45,X46)|k1_funct_1(X46,esk7_2(X45,X46))!=esk6_2(X45,X46)|k2_relat_1(X45)!=k1_relat_1(X46)|k1_relat_1(X45)!=k2_relat_1(X46)|~v2_funct_1(X45)|X46=k2_funct_1(X45)|(~v1_relat_1(X46)|~v1_funct_1(X46))|(~v1_relat_1(X45)|~v1_funct_1(X45)))&(k1_funct_1(X45,esk6_2(X45,X46))=esk7_2(X45,X46)|k1_funct_1(X46,esk7_2(X45,X46))=esk6_2(X45,X46)|k2_relat_1(X45)!=k1_relat_1(X46)|k1_relat_1(X45)!=k2_relat_1(X46)|~v2_funct_1(X45)|X46=k2_funct_1(X45)|(~v1_relat_1(X46)|~v1_funct_1(X46))|(~v1_relat_1(X45)|~v1_funct_1(X45))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t60_funct_1])])])])])).
cnf(c_0_97, negated_conjecture, (k2_relat_1(esk4_0)=esk1_0|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_relat_1(esk4_0)))), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_98, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k2_relat_1(esk4_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_33]), c_0_47]), c_0_48]), c_0_49]), c_0_50])])).
cnf(c_0_99, plain, (k1_funct_1(X1,esk6_2(X1,X2))=esk7_2(X1,X2)|k1_funct_1(X2,esk7_2(X1,X2))=esk6_2(X1,X2)|X2=k2_funct_1(X1)|k2_relat_1(X1)!=k1_relat_1(X2)|k1_relat_1(X1)!=k2_relat_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_100, negated_conjecture, (k2_relat_1(esk4_0)=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_98])])).
cnf(c_0_101, plain, (r2_hidden(esk7_2(X1,X2),k1_relat_1(X2))|X2=k2_funct_1(X1)|k2_relat_1(X1)!=k1_relat_1(X2)|k1_relat_1(X1)!=k2_relat_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_102, plain, (X2=k2_funct_1(X1)|k1_funct_1(X1,esk6_2(X1,X2))!=esk7_2(X1,X2)|k1_funct_1(X2,esk7_2(X1,X2))!=esk6_2(X1,X2)|k2_relat_1(X1)!=k1_relat_1(X2)|k1_relat_1(X1)!=k2_relat_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_103, negated_conjecture, (k1_funct_1(esk3_0,X1)=X2|~r2_hidden(X2,esk2_0)|k1_funct_1(esk4_0,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_104, negated_conjecture, (k1_funct_1(esk4_0,esk7_2(X1,esk4_0))=esk6_2(X1,esk4_0)|k1_funct_1(X1,esk6_2(X1,esk4_0))=esk7_2(X1,esk4_0)|k2_funct_1(X1)=esk4_0|k1_relat_1(X1)!=esk1_0|k2_relat_1(X1)!=esk2_0|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_59]), c_0_57]), c_0_58])])).
cnf(c_0_105, negated_conjecture, (esk4_0!=k2_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_106, negated_conjecture, (k2_funct_1(X1)=esk4_0|r2_hidden(esk7_2(X1,esk4_0),esk2_0)|k1_relat_1(X1)!=esk1_0|k2_relat_1(X1)!=esk2_0|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_100]), c_0_59]), c_0_59]), c_0_57]), c_0_58])])).
cnf(c_0_107, plain, (r2_hidden(esk6_2(X1,X2),k1_relat_1(X1))|X2=k2_funct_1(X1)|k2_relat_1(X1)!=k1_relat_1(X2)|k1_relat_1(X1)!=k2_relat_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_108, negated_conjecture, (k2_funct_1(X1)=esk4_0|k1_funct_1(esk4_0,esk7_2(X1,esk4_0))!=esk6_2(X1,esk4_0)|k1_funct_1(X1,esk6_2(X1,esk4_0))!=esk7_2(X1,esk4_0)|k1_relat_1(X1)!=esk1_0|k2_relat_1(X1)!=esk2_0|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_100]), c_0_59]), c_0_57]), c_0_58])])).
cnf(c_0_109, negated_conjecture, (k1_funct_1(esk3_0,k1_funct_1(esk4_0,X1))=X1|~r2_hidden(X1,esk2_0)), inference(er,[status(thm)],[c_0_103])).
cnf(c_0_110, negated_conjecture, (k1_funct_1(esk3_0,esk6_2(esk3_0,esk4_0))=esk7_2(esk3_0,esk4_0)|k1_funct_1(esk4_0,esk7_2(esk3_0,esk4_0))=esk6_2(esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_47]), c_0_64]), c_0_48]), c_0_49]), c_0_50])]), c_0_105])).
cnf(c_0_111, negated_conjecture, (r2_hidden(esk7_2(esk3_0,esk4_0),esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_47]), c_0_64]), c_0_48]), c_0_49]), c_0_50])]), c_0_105])).
cnf(c_0_112, negated_conjecture, (k2_funct_1(X1)=esk4_0|r2_hidden(esk6_2(X1,esk4_0),k1_relat_1(X1))|k1_relat_1(X1)!=esk1_0|k2_relat_1(X1)!=esk2_0|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_100]), c_0_59]), c_0_57]), c_0_58])])).
cnf(c_0_113, negated_conjecture, (k1_funct_1(esk4_0,esk7_2(esk3_0,esk4_0))!=esk6_2(esk3_0,esk4_0)|k1_funct_1(esk3_0,esk6_2(esk3_0,esk4_0))!=esk7_2(esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_47]), c_0_64]), c_0_48]), c_0_49]), c_0_50])]), c_0_105])).
cnf(c_0_114, negated_conjecture, (k1_funct_1(esk3_0,esk6_2(esk3_0,esk4_0))=esk7_2(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_111])])).
cnf(c_0_115, negated_conjecture, (r2_hidden(esk6_2(esk3_0,esk4_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_47]), c_0_64]), c_0_48]), c_0_49]), c_0_50])]), c_0_105])).
cnf(c_0_116, negated_conjecture, (k1_funct_1(esk4_0,esk7_2(esk3_0,esk4_0))!=esk6_2(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_113, c_0_114])])).
cnf(c_0_117, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_114]), c_0_115])]), c_0_116]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 118
# Proof object clause steps            : 85
# Proof object formula steps           : 33
# Proof object conjectures             : 60
# Proof object clause conjectures      : 57
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 36
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 46
# Proof object simplifying inferences  : 129
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 52
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 49
# Processed clauses                    : 915
# ...of these trivial                  : 2
# ...subsumed                          : 370
# ...remaining for further processing  : 542
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 14
# Backward-rewritten                   : 52
# Generated clauses                    : 1434
# ...of the previous two non-trivial   : 1306
# Contextual simplify-reflections      : 213
# Paramodulations                      : 1401
# Factorizations                       : 0
# Equation resolutions                 : 33
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 426
#    Positive orientable unit clauses  : 34
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 388
# Current number of unprocessed clauses: 465
# ...number of literals in the above   : 3736
# Current number of archived formulas  : 0
# Current number of archived clauses   : 114
# Clause-clause subsumption calls (NU) : 33850
# Rec. Clause-clause subsumption calls : 8378
# Non-unit clause-clause subsumptions  : 566
# Unit Clause-clause subsumption calls : 322
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 31
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 45772

# -------------------------------------------------
# User time                : 0.050 s
# System time              : 0.003 s
# Total time               : 0.053 s
# Maximum resident set size: 28748 pages
