# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/11614c550efe4e4faab3bd2a2e27b6f3.E_input', d4_xboole_0)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/11614c550efe4e4faab3bd2a2e27b6f3.E_input', t12_setfam_1)).
fof(redefinition_k8_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k8_subset_1(X1,X2,X3)=k3_xboole_0(X2,X3)), file('data/MPTP/proofs/11614c550efe4e4faab3bd2a2e27b6f3.E_input', redefinition_k8_subset_1)).
fof(t39_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X3,u1_struct_0(k9_yellow_6(X1,X2)))<=>(r2_hidden(X2,X3)&v3_pre_topc(X3,X1)))))), file('data/MPTP/proofs/11614c550efe4e4faab3bd2a2e27b6f3.E_input', t39_yellow_6)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/11614c550efe4e4faab3bd2a2e27b6f3.E_input', t4_subset)).
fof(dt_k8_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>m1_subset_1(k8_subset_1(X1,X2,X3),k1_zfmisc_1(X1))), file('data/MPTP/proofs/11614c550efe4e4faab3bd2a2e27b6f3.E_input', dt_k8_subset_1)).
fof(t22_waybel_9, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(k9_yellow_6(X1,X2)))=>![X4]:(m1_subset_1(X4,u1_struct_0(k9_yellow_6(X1,X2)))=>m1_subset_1(k3_xboole_0(X3,X4),u1_struct_0(k9_yellow_6(X1,X2))))))), file('data/MPTP/proofs/11614c550efe4e4faab3bd2a2e27b6f3.E_input', t22_waybel_9)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/11614c550efe4e4faab3bd2a2e27b6f3.E_input', t1_subset)).
fof(fc6_tops_1, axiom, ![X1, X2, X3]:(((v2_pre_topc(X1)&l1_pre_topc(X1))&((v3_pre_topc(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(v3_pre_topc(X3,X1)&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))))=>v3_pre_topc(k3_xboole_0(X2,X3),X1)), file('data/MPTP/proofs/11614c550efe4e4faab3bd2a2e27b6f3.E_input', fc6_tops_1)).
fof(t38_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(k9_yellow_6(X1,X2)))=>?[X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))&(X4=X3&(r2_hidden(X2,X4)&v3_pre_topc(X4,X1))))))), file('data/MPTP/proofs/11614c550efe4e4faab3bd2a2e27b6f3.E_input', t38_yellow_6)).
fof(c_0_10, plain, ![X23, X24, X25, X26, X27, X28, X29, X30]:((((r2_hidden(X26,X23)|~r2_hidden(X26,X25)|X25!=k3_xboole_0(X23,X24))&(r2_hidden(X26,X24)|~r2_hidden(X26,X25)|X25!=k3_xboole_0(X23,X24)))&(~r2_hidden(X27,X23)|~r2_hidden(X27,X24)|r2_hidden(X27,X25)|X25!=k3_xboole_0(X23,X24)))&((~r2_hidden(esk5_3(X28,X29,X30),X30)|(~r2_hidden(esk5_3(X28,X29,X30),X28)|~r2_hidden(esk5_3(X28,X29,X30),X29))|X30=k3_xboole_0(X28,X29))&((r2_hidden(esk5_3(X28,X29,X30),X28)|r2_hidden(esk5_3(X28,X29,X30),X30)|X30=k3_xboole_0(X28,X29))&(r2_hidden(esk5_3(X28,X29,X30),X29)|r2_hidden(esk5_3(X28,X29,X30),X30)|X30=k3_xboole_0(X28,X29))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_11, plain, ![X9, X10]:k1_setfam_1(k2_tarski(X9,X10))=k3_xboole_0(X9,X10), inference(variable_rename,[status(thm)],[t12_setfam_1])).
cnf(c_0_12, plain, (r2_hidden(esk5_3(X1,X2,X3),X2)|r2_hidden(esk5_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_13, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_14, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, plain, ![X11, X12, X13]:(~m1_subset_1(X12,k1_zfmisc_1(X11))|k8_subset_1(X11,X12,X13)=k3_xboole_0(X12,X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_subset_1])])).
cnf(c_0_16, plain, (X3=k3_xboole_0(X1,X2)|~r2_hidden(esk5_3(X1,X2,X3),X3)|~r2_hidden(esk5_3(X1,X2,X3),X1)|~r2_hidden(esk5_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, plain, (X3=k1_setfam_1(k2_tarski(X1,X2))|r2_hidden(esk5_3(X1,X2,X3),X3)|r2_hidden(esk5_3(X1,X2,X3),X2)), inference(rw,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_18, plain, (r2_hidden(X1,X2)|X3!=k1_setfam_1(k2_tarski(X4,X2))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_14, c_0_13])).
fof(c_0_19, plain, ![X17, X18, X19]:(((r2_hidden(X18,X19)|~r2_hidden(X19,u1_struct_0(k9_yellow_6(X17,X18)))|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X17)))|~m1_subset_1(X18,u1_struct_0(X17))|(v2_struct_0(X17)|(~v2_pre_topc(X17)|~l1_pre_topc(X17))))&(v3_pre_topc(X19,X17)|~r2_hidden(X19,u1_struct_0(k9_yellow_6(X17,X18)))|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X17)))|~m1_subset_1(X18,u1_struct_0(X17))|(v2_struct_0(X17)|(~v2_pre_topc(X17)|~l1_pre_topc(X17)))))&(~r2_hidden(X18,X19)|~v3_pre_topc(X19,X17)|r2_hidden(X19,u1_struct_0(k9_yellow_6(X17,X18)))|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X17)))|~m1_subset_1(X18,u1_struct_0(X17))|(v2_struct_0(X17)|(~v2_pre_topc(X17)|~l1_pre_topc(X17))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t39_yellow_6])])])])])).
fof(c_0_20, plain, ![X14, X15, X16]:(~r2_hidden(X14,X15)|~m1_subset_1(X15,k1_zfmisc_1(X16))|m1_subset_1(X14,X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_21, plain, ![X32, X33, X34]:(~m1_subset_1(X33,k1_zfmisc_1(X32))|m1_subset_1(k8_subset_1(X32,X33,X34),k1_zfmisc_1(X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_subset_1])])).
cnf(c_0_22, plain, (k8_subset_1(X2,X1,X3)=k3_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (X3=k1_setfam_1(k2_tarski(X1,X2))|~r2_hidden(esk5_3(X1,X2,X3),X3)|~r2_hidden(esk5_3(X1,X2,X3),X2)|~r2_hidden(esk5_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_16, c_0_13])).
cnf(c_0_24, plain, (k1_setfam_1(k2_tarski(X1,X2))=X2|r2_hidden(esk5_3(X1,X2,X2),X2)), inference(ef,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_setfam_1(k2_tarski(X3,X2)))), inference(er,[status(thm)],[c_0_18])).
fof(c_0_26, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(k9_yellow_6(X1,X2)))=>![X4]:(m1_subset_1(X4,u1_struct_0(k9_yellow_6(X1,X2)))=>m1_subset_1(k3_xboole_0(X3,X4),u1_struct_0(k9_yellow_6(X1,X2)))))))), inference(assume_negation,[status(cth)],[t22_waybel_9])).
fof(c_0_27, plain, ![X39, X40]:(~r2_hidden(X39,X40)|m1_subset_1(X39,X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_28, plain, (r2_hidden(X2,u1_struct_0(k9_yellow_6(X3,X1)))|v2_struct_0(X3)|~r2_hidden(X1,X2)|~v3_pre_topc(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_30, plain, ![X20, X21, X22]:(~v2_pre_topc(X20)|~l1_pre_topc(X20)|(~v3_pre_topc(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|(~v3_pre_topc(X22,X20)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))))|v3_pre_topc(k3_xboole_0(X21,X22),X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc6_tops_1])])).
cnf(c_0_31, plain, (m1_subset_1(k8_subset_1(X2,X1,X3),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_32, plain, (k8_subset_1(X2,X1,X3)=k1_setfam_1(k2_tarski(X1,X3))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_22, c_0_13])).
cnf(c_0_33, plain, (k1_setfam_1(k2_tarski(X1,X2))=X2|~r2_hidden(esk5_3(X1,X2,X2),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_24])).
cnf(c_0_34, plain, (k1_setfam_1(k2_tarski(X1,k1_setfam_1(k2_tarski(X2,X3))))=k1_setfam_1(k2_tarski(X2,X3))|r2_hidden(esk5_3(X1,k1_setfam_1(k2_tarski(X2,X3)),k1_setfam_1(k2_tarski(X2,X3))),X3)), inference(spm,[status(thm)],[c_0_25, c_0_24])).
fof(c_0_35, plain, ![X35, X36, X37]:((m1_subset_1(esk6_3(X35,X36,X37),k1_zfmisc_1(u1_struct_0(X35)))|~m1_subset_1(X37,u1_struct_0(k9_yellow_6(X35,X36)))|~m1_subset_1(X36,u1_struct_0(X35))|(v2_struct_0(X35)|(~v2_pre_topc(X35)|~l1_pre_topc(X35))))&((esk6_3(X35,X36,X37)=X37|~m1_subset_1(X37,u1_struct_0(k9_yellow_6(X35,X36)))|~m1_subset_1(X36,u1_struct_0(X35))|(v2_struct_0(X35)|(~v2_pre_topc(X35)|~l1_pre_topc(X35))))&((r2_hidden(X36,esk6_3(X35,X36,X37))|~m1_subset_1(X37,u1_struct_0(k9_yellow_6(X35,X36)))|~m1_subset_1(X36,u1_struct_0(X35))|(v2_struct_0(X35)|(~v2_pre_topc(X35)|~l1_pre_topc(X35))))&(v3_pre_topc(esk6_3(X35,X36,X37),X35)|~m1_subset_1(X37,u1_struct_0(k9_yellow_6(X35,X36)))|~m1_subset_1(X36,u1_struct_0(X35))|(v2_struct_0(X35)|(~v2_pre_topc(X35)|~l1_pre_topc(X35))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t38_yellow_6])])])])])])).
fof(c_0_36, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))&(m1_subset_1(esk4_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))&~m1_subset_1(k3_xboole_0(esk3_0,esk4_0),u1_struct_0(k9_yellow_6(esk1_0,esk2_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_26])])])])).
cnf(c_0_37, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_38, plain, (r2_hidden(X1,u1_struct_0(k9_yellow_6(X2,X3)))|v2_struct_0(X2)|~v3_pre_topc(X1,X2)|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_39, plain, (v3_pre_topc(k3_xboole_0(X2,X3),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_pre_topc(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40, plain, (m1_subset_1(k1_setfam_1(k2_tarski(X1,X2)),k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_41, plain, (k1_setfam_1(k2_tarski(X1,k1_setfam_1(k2_tarski(X2,X1))))=k1_setfam_1(k2_tarski(X2,X1))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_42, plain, (v3_pre_topc(esk6_3(X1,X2,X3),X1)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(k9_yellow_6(X1,X2)))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_43, plain, (esk6_3(X1,X2,X3)=X3|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(k9_yellow_6(X1,X2)))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_44, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(k9_yellow_6(X1,X2)))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_45, negated_conjecture, (~m1_subset_1(k3_xboole_0(esk3_0,esk4_0),u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_46, plain, (m1_subset_1(X1,u1_struct_0(k9_yellow_6(X2,X3)))|v2_struct_0(X2)|~v3_pre_topc(X1,X2)|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_47, plain, (v3_pre_topc(k1_setfam_1(k2_tarski(X2,X3)),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v3_pre_topc(X3,X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(rw,[status(thm)],[c_0_39, c_0_13])).
cnf(c_0_48, plain, (m1_subset_1(k1_setfam_1(k2_tarski(X1,X2)),k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_49, plain, (v3_pre_topc(X1,X2)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(k9_yellow_6(X2,X3)))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_50, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_51, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_52, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_53, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_54, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_55, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_56, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(k9_yellow_6(X2,X3)))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_44, c_0_43])).
cnf(c_0_57, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X4!=k3_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_58, plain, (r2_hidden(X1,esk6_3(X2,X1,X3))|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(k9_yellow_6(X2,X1)))|~m1_subset_1(X1,u1_struct_0(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_59, negated_conjecture, (~m1_subset_1(k1_setfam_1(k2_tarski(esk3_0,esk4_0)),u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(rw,[status(thm)],[c_0_45, c_0_13])).
cnf(c_0_60, plain, (m1_subset_1(k1_setfam_1(k2_tarski(X1,X2)),u1_struct_0(k9_yellow_6(X3,X4)))|v2_struct_0(X3)|~v3_pre_topc(X2,X3)|~v3_pre_topc(X1,X3)|~r2_hidden(X4,k1_setfam_1(k2_tarski(X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3)|~v2_pre_topc(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])).
cnf(c_0_61, negated_conjecture, (v3_pre_topc(esk4_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51]), c_0_52]), c_0_53])]), c_0_54])).
cnf(c_0_62, negated_conjecture, (v3_pre_topc(esk3_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_55]), c_0_51]), c_0_52]), c_0_53])]), c_0_54])).
cnf(c_0_63, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_50]), c_0_51]), c_0_52]), c_0_53])]), c_0_54])).
cnf(c_0_64, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_55]), c_0_51]), c_0_52]), c_0_53])]), c_0_54])).
cnf(c_0_65, plain, (r2_hidden(X1,X4)|X4!=k1_setfam_1(k2_tarski(X2,X3))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_57, c_0_13])).
cnf(c_0_66, plain, (r2_hidden(X1,X2)|v2_struct_0(X3)|~m1_subset_1(X2,u1_struct_0(k9_yellow_6(X3,X1)))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_pre_topc(X3)|~v2_pre_topc(X3)), inference(spm,[status(thm)],[c_0_58, c_0_43])).
cnf(c_0_67, negated_conjecture, (~r2_hidden(esk2_0,k1_setfam_1(k2_tarski(esk3_0,esk4_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61]), c_0_62]), c_0_63]), c_0_64]), c_0_52]), c_0_53])]), c_0_54])).
cnf(c_0_68, plain, (r2_hidden(X1,k1_setfam_1(k2_tarski(X2,X3)))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_65])).
cnf(c_0_69, negated_conjecture, (r2_hidden(esk2_0,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_50]), c_0_51]), c_0_52]), c_0_53])]), c_0_54])).
cnf(c_0_70, negated_conjecture, (r2_hidden(esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_55]), c_0_51]), c_0_52]), c_0_53])]), c_0_54])).
cnf(c_0_71, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_69]), c_0_70])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 72
# Proof object clause steps            : 51
# Proof object formula steps           : 21
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 53
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 25
# Processed clauses                    : 5332
# ...of these trivial                  : 0
# ...subsumed                          : 3248
# ...remaining for further processing  : 2084
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 58
# Backward-rewritten                   : 1
# Generated clauses                    : 381150
# ...of the previous two non-trivial   : 374646
# Contextual simplify-reflections      : 60
# Paramodulations                      : 380847
# Factorizations                       : 300
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1997
#    Positive orientable unit clauses  : 22
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 1972
# Current number of unprocessed clauses: 369299
# ...number of literals in the above   : 1803679
# Current number of archived formulas  : 0
# Current number of archived clauses   : 85
# Clause-clause subsumption calls (NU) : 553446
# Rec. Clause-clause subsumption calls : 196194
# Non-unit clause-clause subsumptions  : 3366
# Unit Clause-clause subsumption calls : 1198
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 42
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 12316168

# -------------------------------------------------
# User time                : 5.199 s
# System time              : 0.120 s
# Total time               : 5.319 s
# Maximum resident set size: 287536 pages
