# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C47_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.007 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t58_yellow_0, conjecture, ![X1]:(l1_orders_2(X1)=>![X2]:((v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1))=>![X3]:((v4_yellow_0(X3,X1)&m1_yellow_0(X3,X1))=>(u1_struct_0(X2)=u1_struct_0(X3)=>g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=g1_orders_2(u1_struct_0(X3),u1_orders_2(X3)))))), file('data/MPTP/proofs/1170d3f9a4ec4e459dfc6eaca8315bfe.E_input', t58_yellow_0)).
fof(d14_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_yellow_0(X2,X1)=>(v4_yellow_0(X2,X1)<=>u1_orders_2(X2)=k1_toler_1(u1_orders_2(X1),u1_struct_0(X2))))), file('data/MPTP/proofs/1170d3f9a4ec4e459dfc6eaca8315bfe.E_input', d14_yellow_0)).
fof(c_0_2, negated_conjecture, ~(![X1]:(l1_orders_2(X1)=>![X2]:((v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1))=>![X3]:((v4_yellow_0(X3,X1)&m1_yellow_0(X3,X1))=>(u1_struct_0(X2)=u1_struct_0(X3)=>g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))))))), inference(assume_negation,[status(cth)],[t58_yellow_0])).
fof(c_0_3, plain, ![X7, X8]:((~v4_yellow_0(X8,X7)|u1_orders_2(X8)=k1_toler_1(u1_orders_2(X7),u1_struct_0(X8))|~m1_yellow_0(X8,X7)|~l1_orders_2(X7))&(u1_orders_2(X8)!=k1_toler_1(u1_orders_2(X7),u1_struct_0(X8))|v4_yellow_0(X8,X7)|~m1_yellow_0(X8,X7)|~l1_orders_2(X7))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d14_yellow_0])])])])).
fof(c_0_4, negated_conjecture, (l1_orders_2(esk1_0)&((v4_yellow_0(esk2_0,esk1_0)&m1_yellow_0(esk2_0,esk1_0))&((v4_yellow_0(esk3_0,esk1_0)&m1_yellow_0(esk3_0,esk1_0))&(u1_struct_0(esk2_0)=u1_struct_0(esk3_0)&g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(u1_struct_0(esk3_0),u1_orders_2(esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
cnf(c_0_5, plain, (u1_orders_2(X1)=k1_toler_1(u1_orders_2(X2),u1_struct_0(X1))|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_6, negated_conjecture, (u1_struct_0(esk2_0)=u1_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, negated_conjecture, (k1_toler_1(u1_orders_2(X1),u1_struct_0(esk3_0))=u1_orders_2(esk2_0)|~m1_yellow_0(esk2_0,X1)|~v4_yellow_0(esk2_0,X1)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_8, negated_conjecture, (g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(u1_struct_0(esk3_0),u1_orders_2(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9, negated_conjecture, (u1_orders_2(esk2_0)=u1_orders_2(esk3_0)|~m1_yellow_0(esk3_0,X1)|~m1_yellow_0(esk2_0,X1)|~v4_yellow_0(esk3_0,X1)|~v4_yellow_0(esk2_0,X1)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_5, c_0_7])).
cnf(c_0_10, negated_conjecture, (m1_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_11, negated_conjecture, (m1_yellow_0(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_12, negated_conjecture, (v4_yellow_0(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_13, negated_conjecture, (v4_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_14, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_15, negated_conjecture, (g1_orders_2(u1_struct_0(esk3_0),u1_orders_2(esk2_0))!=g1_orders_2(u1_struct_0(esk3_0),u1_orders_2(esk3_0))), inference(rw,[status(thm)],[c_0_8, c_0_6])).
cnf(c_0_16, negated_conjecture, (u1_orders_2(esk2_0)=u1_orders_2(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11]), c_0_12]), c_0_13]), c_0_14])])).
cnf(c_0_17, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_16])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 18
# Proof object clause steps            : 13
# Proof object formula steps           : 5
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 2
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 2
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 9
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 9
# Processed clauses                    : 24
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 24
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 6
# Generated clauses                    : 12
# ...of the previous two non-trivial   : 10
# Contextual simplify-reflections      : 0
# Paramodulations                      : 12
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 9
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 2
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 15
# Clause-clause subsumption calls (NU) : 2
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 803

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.000 s
# Total time               : 0.008 s
# Maximum resident set size: 28684 pages
