# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', t70_enumset1)).
fof(t31_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>r1_tarski(k3_relat_1(X1),k3_relat_1(X2))))), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', t31_relat_1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', t12_setfam_1)).
fof(t43_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3)), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', t43_xboole_1)).
fof(t10_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(X1,k2_xboole_0(X3,X2))), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', t10_xboole_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', d3_tarski)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', t36_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', t48_xboole_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', commutativity_k2_tarski)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', d10_xboole_0)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', t2_xboole_1)).
fof(t11_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k2_xboole_0(X1,X2),X3)=>r1_tarski(X1,X3)), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', t11_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', t3_boole)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', t4_xboole_0)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', symmetry_r1_xboole_0)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', t65_xboole_1)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', d4_relat_1)).
fof(cc1_relat_1, axiom, ![X1]:(v1_xboole_0(X1)=>v1_relat_1(X1)), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', cc1_relat_1)).
fof(t8_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X3,X2))=>r1_tarski(k2_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', t8_xboole_1)).
fof(d6_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k3_relat_1(X1)=k2_xboole_0(k1_relat_1(X1),k2_relat_1(X1))), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', d6_relat_1)).
fof(t15_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>r1_tarski(k6_subset_1(k1_relat_1(X1),k1_relat_1(X2)),k1_relat_1(k6_subset_1(X1,X2))))), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', t15_relat_1)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', redefinition_k6_subset_1)).
fof(t19_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>~((r2_hidden(X1,k2_relat_1(X2))&![X3]:~(r2_hidden(X3,k1_relat_1(X2)))))), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', t19_relat_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', fc1_xboole_0)).
fof(t44_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k4_xboole_0(X1,X2),X3)=>r1_tarski(X1,k2_xboole_0(X2,X3))), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', t44_xboole_1)).
fof(t28_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>r1_tarski(k6_subset_1(k2_relat_1(X1),k2_relat_1(X2)),k2_relat_1(k6_subset_1(X1,X2))))), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', t28_relat_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/1199720d59784b628ce2c702c2655a75.E_input', l32_xboole_1)).
fof(c_0_28, plain, ![X5, X6]:k3_tarski(k2_tarski(X5,X6))=k2_xboole_0(X5,X6), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_29, plain, ![X7, X8]:k1_enumset1(X7,X7,X8)=k2_tarski(X7,X8), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_30, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>r1_tarski(k3_relat_1(X1),k3_relat_1(X2)))))), inference(assume_negation,[status(cth)],[t31_relat_1])).
fof(c_0_31, plain, ![X15, X16]:k1_setfam_1(k2_tarski(X15,X16))=k3_xboole_0(X15,X16), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_32, plain, ![X17, X18, X19]:(~r1_tarski(X17,k2_xboole_0(X18,X19))|r1_tarski(k4_xboole_0(X17,X18),X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
cnf(c_0_33, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_35, plain, ![X20, X21, X22]:(~r1_tarski(X20,X21)|r1_tarski(X20,k2_xboole_0(X22,X21))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_xboole_1])])).
fof(c_0_36, plain, ![X23, X24, X25, X26, X27]:((~r1_tarski(X23,X24)|(~r2_hidden(X25,X23)|r2_hidden(X25,X24)))&((r2_hidden(esk3_2(X26,X27),X26)|r1_tarski(X26,X27))&(~r2_hidden(esk3_2(X26,X27),X27)|r1_tarski(X26,X27)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_37, negated_conjecture, (v1_relat_1(esk1_0)&(v1_relat_1(esk2_0)&(r1_tarski(esk1_0,esk2_0)&~r1_tarski(k3_relat_1(esk1_0),k3_relat_1(esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])).
fof(c_0_38, plain, ![X11, X12]:r1_tarski(k4_xboole_0(X11,X12),X11), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_39, plain, ![X29, X30]:k4_xboole_0(X29,k4_xboole_0(X29,X30))=k3_xboole_0(X29,X30), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_40, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_41, plain, ![X31, X32]:k2_tarski(X31,X32)=k2_tarski(X32,X31), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
fof(c_0_42, plain, ![X33, X34]:(((r1_tarski(X33,X34)|X33!=X34)&(r1_tarski(X34,X33)|X33!=X34))&(~r1_tarski(X33,X34)|~r1_tarski(X34,X33)|X33=X34)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_43, plain, ![X35]:r1_tarski(k1_xboole_0,X35), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_44, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_45, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_46, plain, (r1_tarski(X1,k2_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_47, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_48, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_49, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_40, c_0_34])).
cnf(c_0_52, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_53, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_54, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_55, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k3_tarski(k1_enumset1(X2,X2,X3)))), inference(rw,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_56, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X3,X3,X2)))|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_46, c_0_45])).
cnf(c_0_57, plain, (r1_tarski(X1,X2)|~r2_hidden(esk3_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_58, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_59, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_47, c_0_49])).
cnf(c_0_60, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_61, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_setfam_1(k1_enumset1(X1,X1,X2))), inference(rw,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_62, plain, (k1_enumset1(X1,X1,X2)=k1_enumset1(X2,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_34]), c_0_34])).
cnf(c_0_63, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_64, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,X3)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
fof(c_0_65, plain, ![X37, X38, X39]:(~r1_tarski(k2_xboole_0(X37,X38),X39)|r1_tarski(X37,X39)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_xboole_1])])).
cnf(c_0_66, negated_conjecture, (r1_tarski(X1,esk2_0)|~r2_hidden(esk3_2(X1,esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_67, plain, (r2_hidden(esk3_2(k4_xboole_0(X1,X2),X3),X1)|r1_tarski(k4_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_68, plain, (k4_xboole_0(X1,X2)=X1|~r1_tarski(X1,k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_53, c_0_49])).
cnf(c_0_69, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_61])).
cnf(c_0_70, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
fof(c_0_71, plain, ![X36]:k4_xboole_0(X36,k1_xboole_0)=X36, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_72, plain, ![X40, X41, X43, X44, X45]:((r1_xboole_0(X40,X41)|r2_hidden(esk4_2(X40,X41),k3_xboole_0(X40,X41)))&(~r2_hidden(X45,k3_xboole_0(X43,X44))|~r1_xboole_0(X43,X44))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
cnf(c_0_73, plain, (r1_tarski(X1,X3)|~r1_tarski(k2_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_74, negated_conjecture, (r1_tarski(k4_xboole_0(esk1_0,X1),esk2_0)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_75, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=X2|~r1_tarski(X2,k4_xboole_0(X1,k4_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_76, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_xboole_0|~r1_tarski(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_70, c_0_69])).
cnf(c_0_77, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_78, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_79, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_72])).
fof(c_0_80, plain, ![X46, X47]:(~r1_xboole_0(X46,X47)|r1_xboole_0(X47,X46)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
fof(c_0_81, plain, ![X48]:r1_xboole_0(X48,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_82, plain, (r1_tarski(X1,X3)|~r1_tarski(k3_tarski(k1_enumset1(X1,X1,X2)),X3)), inference(rw,[status(thm)],[c_0_73, c_0_45])).
cnf(c_0_83, negated_conjecture, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,esk1_0)),esk2_0)), inference(spm,[status(thm)],[c_0_74, c_0_69])).
cnf(c_0_84, plain, (k4_xboole_0(X1,X2)=X1|~r1_tarski(X2,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77]), c_0_77]), c_0_49])])).
cnf(c_0_85, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_78])).
cnf(c_0_86, plain, (~r1_xboole_0(X2,X3)|~r2_hidden(X1,k1_setfam_1(k1_enumset1(X2,X2,X3)))), inference(rw,[status(thm)],[c_0_79, c_0_51])).
cnf(c_0_87, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_88, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_81])).
fof(c_0_89, plain, ![X49, X50, X51, X53, X54, X55, X56, X58]:(((~r2_hidden(X51,X50)|r2_hidden(k4_tarski(X51,esk5_3(X49,X50,X51)),X49)|X50!=k1_relat_1(X49))&(~r2_hidden(k4_tarski(X53,X54),X49)|r2_hidden(X53,X50)|X50!=k1_relat_1(X49)))&((~r2_hidden(esk6_2(X55,X56),X56)|~r2_hidden(k4_tarski(esk6_2(X55,X56),X58),X55)|X56=k1_relat_1(X55))&(r2_hidden(esk6_2(X55,X56),X56)|r2_hidden(k4_tarski(esk6_2(X55,X56),esk7_2(X55,X56)),X55)|X56=k1_relat_1(X55)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
cnf(c_0_90, plain, (r1_tarski(X1,X2)|~r1_tarski(k3_tarski(k1_enumset1(X3,X3,X1)),X2)), inference(spm,[status(thm)],[c_0_82, c_0_62])).
cnf(c_0_91, negated_conjecture, (r1_tarski(X1,esk2_0)|~r1_tarski(k4_xboole_0(X1,esk1_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_92, plain, (r1_tarski(k4_xboole_0(k3_tarski(k1_enumset1(X1,X1,X2)),X1),X2)), inference(spm,[status(thm)],[c_0_55, c_0_85])).
cnf(c_0_93, plain, (~r1_xboole_0(X1,X2)|~r2_hidden(X3,k4_xboole_0(X1,k4_xboole_0(X1,X2)))), inference(rw,[status(thm)],[c_0_86, c_0_61])).
cnf(c_0_94, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_63, c_0_49])).
cnf(c_0_95, plain, (r1_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_96, plain, (r2_hidden(k4_tarski(X1,esk5_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_89])).
fof(c_0_97, plain, ![X60]:(~v1_xboole_0(X60)|v1_relat_1(X60)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relat_1])])).
fof(c_0_98, plain, ![X61, X62, X63]:(~r1_tarski(X61,X62)|~r1_tarski(X63,X62)|r1_tarski(k2_xboole_0(X61,X63),X62)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_xboole_1])])).
fof(c_0_99, plain, ![X64]:(~v1_relat_1(X64)|k3_relat_1(X64)=k2_xboole_0(k1_relat_1(X64),k2_relat_1(X64))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_relat_1])])).
fof(c_0_100, plain, ![X71, X72]:(~v1_relat_1(X71)|(~v1_relat_1(X72)|r1_tarski(k6_subset_1(k1_relat_1(X71),k1_relat_1(X72)),k1_relat_1(k6_subset_1(X71,X72))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_relat_1])])])).
fof(c_0_101, plain, ![X13, X14]:k6_subset_1(X13,X14)=k4_xboole_0(X13,X14), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
cnf(c_0_102, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X2,X2,X3)))|~r1_tarski(k3_tarski(k1_enumset1(X4,X4,X1)),X3)), inference(spm,[status(thm)],[c_0_90, c_0_56])).
cnf(c_0_103, negated_conjecture, (r1_tarski(k3_tarski(k1_enumset1(k1_xboole_0,k1_xboole_0,esk1_0)),esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_62])).
cnf(c_0_104, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_95]), c_0_94])])).
cnf(c_0_105, plain, (r2_hidden(k4_tarski(X1,esk5_3(X2,k1_relat_1(X2),X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_96])).
fof(c_0_106, plain, ![X65, X66]:(~v1_relat_1(X66)|(~r2_hidden(X65,k2_relat_1(X66))|r2_hidden(esk8_2(X65,X66),k1_relat_1(X66)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t19_relat_1])])])])).
cnf(c_0_107, plain, (v1_relat_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_108, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_109, plain, (r1_tarski(k2_xboole_0(X1,X3),X2)|~r1_tarski(X1,X2)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_110, plain, (k3_relat_1(X1)=k2_xboole_0(k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_99])).
fof(c_0_111, plain, ![X68, X69, X70]:(~r1_tarski(k4_xboole_0(X68,X69),X70)|r1_tarski(X68,k2_xboole_0(X69,X70))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t44_xboole_1])])).
cnf(c_0_112, plain, (r1_tarski(k6_subset_1(k1_relat_1(X1),k1_relat_1(X2)),k1_relat_1(k6_subset_1(X1,X2)))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_113, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_114, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k3_tarski(k1_enumset1(X3,X3,X2)))), inference(spm,[status(thm)],[c_0_55, c_0_62])).
cnf(c_0_115, negated_conjecture, (r1_tarski(esk1_0,k3_tarski(k1_enumset1(X1,X1,esk2_0)))), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_116, plain, (~r2_hidden(X1,k1_relat_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_104, c_0_105])).
fof(c_0_117, plain, ![X73, X74]:(~v1_relat_1(X73)|(~v1_relat_1(X74)|r1_tarski(k6_subset_1(k2_relat_1(X73),k2_relat_1(X74)),k2_relat_1(k6_subset_1(X73,X74))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_relat_1])])])).
cnf(c_0_118, plain, (r2_hidden(esk8_2(X2,X1),k1_relat_1(X1))|~v1_relat_1(X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_119, plain, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_120, plain, (r1_tarski(k3_tarski(k1_enumset1(X1,X1,X3)),X2)|~r1_tarski(X3,X2)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_109, c_0_45])).
cnf(c_0_121, plain, (k3_relat_1(X1)=k3_tarski(k1_enumset1(k1_relat_1(X1),k1_relat_1(X1),k2_relat_1(X1)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_110, c_0_45])).
cnf(c_0_122, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(k4_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_123, plain, (r1_tarski(k4_xboole_0(k1_relat_1(X1),k1_relat_1(X2)),k1_relat_1(k4_xboole_0(X1,X2)))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_112, c_0_113]), c_0_113])).
cnf(c_0_124, negated_conjecture, (r1_tarski(k4_xboole_0(esk1_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_125, plain, (r1_tarski(k1_relat_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_116, c_0_60])).
cnf(c_0_126, plain, (r1_tarski(k6_subset_1(k2_relat_1(X1),k2_relat_1(X2)),k2_relat_1(k6_subset_1(X1,X2)))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_117])).
cnf(c_0_127, plain, (~r2_hidden(X1,k2_relat_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_118]), c_0_119])])).
cnf(c_0_128, negated_conjecture, (~r1_tarski(k3_relat_1(esk1_0),k3_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_129, plain, (r1_tarski(k3_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X2)|~r1_tarski(k1_relat_1(X1),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_130, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_131, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X2,X2,X3)))|~r1_tarski(k4_xboole_0(X1,X2),X3)), inference(rw,[status(thm)],[c_0_122, c_0_45])).
cnf(c_0_132, plain, (k4_xboole_0(k1_relat_1(X1),k1_relat_1(X2))=k1_relat_1(k4_xboole_0(X1,X2))|~r1_tarski(k1_relat_1(k4_xboole_0(X1,X2)),k4_xboole_0(k1_relat_1(X1),k1_relat_1(X2)))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_53, c_0_123])).
cnf(c_0_133, negated_conjecture, (k4_xboole_0(esk1_0,esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_63, c_0_124])).
cnf(c_0_134, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_63, c_0_125])).
cnf(c_0_135, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_136, plain, (r1_tarski(k4_xboole_0(k2_relat_1(X1),k2_relat_1(X2)),k2_relat_1(k4_xboole_0(X1,X2)))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_126, c_0_113]), c_0_113])).
cnf(c_0_137, plain, (r1_tarski(k2_relat_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_127, c_0_60])).
cnf(c_0_138, negated_conjecture, (~r1_tarski(k2_relat_1(esk1_0),k3_relat_1(esk2_0))|~r1_tarski(k1_relat_1(esk1_0),k3_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129]), c_0_130])])).
cnf(c_0_139, plain, (r1_tarski(X1,k3_relat_1(X2))|~r1_tarski(X1,k2_relat_1(X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_56, c_0_121])).
cnf(c_0_140, plain, (r1_tarski(X1,k3_relat_1(X2))|~r1_tarski(k4_xboole_0(X1,k1_relat_1(X2)),k2_relat_1(X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_131, c_0_121])).
cnf(c_0_141, negated_conjecture, (k4_xboole_0(k1_relat_1(esk1_0),k1_relat_1(esk2_0))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_133]), c_0_134]), c_0_134]), c_0_54]), c_0_135]), c_0_130])])).
fof(c_0_142, plain, ![X75, X76]:((k4_xboole_0(X75,X76)!=k1_xboole_0|r1_tarski(X75,X76))&(~r1_tarski(X75,X76)|k4_xboole_0(X75,X76)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_143, plain, (k4_xboole_0(k2_relat_1(X1),k2_relat_1(X2))=k2_relat_1(k4_xboole_0(X1,X2))|~r1_tarski(k2_relat_1(k4_xboole_0(X1,X2)),k4_xboole_0(k2_relat_1(X1),k2_relat_1(X2)))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_53, c_0_136])).
cnf(c_0_144, plain, (k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_63, c_0_137])).
cnf(c_0_145, negated_conjecture, (~r1_tarski(k1_relat_1(esk1_0),k3_relat_1(esk2_0))|~r1_tarski(k2_relat_1(esk1_0),k2_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_135])])).
cnf(c_0_146, negated_conjecture, (r1_tarski(k1_relat_1(esk1_0),k3_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_141]), c_0_54]), c_0_135])])).
cnf(c_0_147, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_142])).
cnf(c_0_148, negated_conjecture, (k4_xboole_0(k2_relat_1(esk1_0),k2_relat_1(esk2_0))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_133]), c_0_144]), c_0_144]), c_0_54]), c_0_135]), c_0_130])])).
cnf(c_0_149, negated_conjecture, (~r1_tarski(k2_relat_1(esk1_0),k2_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_145, c_0_146])])).
cnf(c_0_150, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_148]), c_0_149]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 151
# Proof object clause steps            : 95
# Proof object formula steps           : 56
# Proof object conjectures             : 23
# Proof object clause conjectures      : 20
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 34
# Proof object initial formulas used   : 28
# Proof object generating inferences   : 44
# Proof object simplifying inferences  : 52
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 28
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 36
# Processed clauses                    : 2454
# ...of these trivial                  : 5
# ...subsumed                          : 1950
# ...remaining for further processing  : 499
# Other redundant clauses eliminated   : 19
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 14
# Generated clauses                    : 12299
# ...of the previous two non-trivial   : 10762
# Contextual simplify-reflections      : 3
# Paramodulations                      : 12276
# Factorizations                       : 4
# Equation resolutions                 : 19
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 442
#    Positive orientable unit clauses  : 37
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 12
#    Non-unit-clauses                  : 391
# Current number of unprocessed clauses: 8321
# ...number of literals in the above   : 28800
# Current number of archived formulas  : 0
# Current number of archived clauses   : 57
# Clause-clause subsumption calls (NU) : 30251
# Rec. Clause-clause subsumption calls : 17877
# Non-unit clause-clause subsumptions  : 898
# Unit Clause-clause subsumption calls : 1372
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 69
# BW rewrite match successes           : 16
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 135964

# -------------------------------------------------
# User time                : 0.095 s
# System time              : 0.007 s
# Total time               : 0.102 s
# Maximum resident set size: 28788 pages
