# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t122_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_pre_topc(X2,X1)=>((v1_tsep_1(X2,X1)&m1_pre_topc(X2,X1))<=>(v1_funct_1(k9_tmap_1(X1,X2))&(v1_funct_2(k9_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))&(v5_pre_topc(k9_tmap_1(X1,X2),X1,k8_tmap_1(X1,X2))&m1_subset_1(k9_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2))))))))))), file('data/MPTP/proofs/11a533fe73c443af9430d72b8ecb8806.E_input', t122_tmap_1)).
fof(d11_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:((v1_pre_topc(X3)&(v2_pre_topc(X3)&l1_pre_topc(X3)))=>(X3=k8_tmap_1(X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(X4=u1_struct_0(X2)=>X3=k6_tmap_1(X1,X4))))))), file('data/MPTP/proofs/11a533fe73c443af9430d72b8ecb8806.E_input', d11_tmap_1)).
fof(fc4_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(~(v2_struct_0(k6_tmap_1(X1,X2)))&(v1_pre_topc(k6_tmap_1(X1,X2))&v2_pre_topc(k6_tmap_1(X1,X2))))), file('data/MPTP/proofs/11a533fe73c443af9430d72b8ecb8806.E_input', fc4_tmap_1)).
fof(t1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/11a533fe73c443af9430d72b8ecb8806.E_input', t1_tsep_1)).
fof(dt_k6_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_pre_topc(k6_tmap_1(X1,X2))&(v2_pre_topc(k6_tmap_1(X1,X2))&l1_pre_topc(k6_tmap_1(X1,X2))))), file('data/MPTP/proofs/11a533fe73c443af9430d72b8ecb8806.E_input', dt_k6_tmap_1)).
fof(t104_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(u1_struct_0(k6_tmap_1(X1,X2))=u1_struct_0(X1)&u1_pre_topc(k6_tmap_1(X1,X2))=k5_tmap_1(X1,X2)))), file('data/MPTP/proofs/11a533fe73c443af9430d72b8ecb8806.E_input', t104_tmap_1)).
fof(t2_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>m1_pre_topc(X1,X1)), file('data/MPTP/proofs/11a533fe73c443af9430d72b8ecb8806.E_input', t2_tsep_1)).
fof(t39_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/11a533fe73c443af9430d72b8ecb8806.E_input', t39_pre_topc)).
fof(d10_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k7_tmap_1(X1,X2)=k6_partfun1(u1_struct_0(X1)))), file('data/MPTP/proofs/11a533fe73c443af9430d72b8ecb8806.E_input', d10_tmap_1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/11a533fe73c443af9430d72b8ecb8806.E_input', fc2_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/11a533fe73c443af9430d72b8ecb8806.E_input', dt_l1_pre_topc)).
fof(dt_k9_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_pre_topc(X2,X1))=>(v1_funct_1(k9_tmap_1(X1,X2))&(v1_funct_2(k9_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))&m1_subset_1(k9_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))))), file('data/MPTP/proofs/11a533fe73c443af9430d72b8ecb8806.E_input', dt_k9_tmap_1)).
fof(d12_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))))=>(X3=k9_tmap_1(X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(X4=u1_struct_0(X2)=>r1_funct_2(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X4)),X3,k7_tmap_1(X1,X4)))))))), file('data/MPTP/proofs/11a533fe73c443af9430d72b8ecb8806.E_input', d12_tmap_1)).
fof(dt_k7_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_funct_1(k7_tmap_1(X1,X2))&(v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))&m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))))), file('data/MPTP/proofs/11a533fe73c443af9430d72b8ecb8806.E_input', dt_k7_tmap_1)).
fof(reflexivity_r1_funct_2, axiom, ![X1, X2, X3, X4, X5, X6]:((~(v1_xboole_0(X2))&(~(v1_xboole_0(X4))&((v1_funct_1(X5)&(v1_funct_2(X5,X1,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&(v1_funct_1(X6)&(v1_funct_2(X6,X3,X4)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))))))=>r1_funct_2(X1,X2,X3,X4,X5,X5)), file('data/MPTP/proofs/11a533fe73c443af9430d72b8ecb8806.E_input', reflexivity_r1_funct_2)).
fof(t113_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>(v1_funct_1(k7_tmap_1(X1,X2))&(v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))&(v5_pre_topc(k7_tmap_1(X1,X2),X1,k6_tmap_1(X1,X2))&m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2))))))))))), file('data/MPTP/proofs/11a533fe73c443af9430d72b8ecb8806.E_input', t113_tmap_1)).
fof(d1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>(v1_tsep_1(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>v3_pre_topc(X3,X1)))))), file('data/MPTP/proofs/11a533fe73c443af9430d72b8ecb8806.E_input', d1_tsep_1)).
fof(c_0_17, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_pre_topc(X2,X1)=>((v1_tsep_1(X2,X1)&m1_pre_topc(X2,X1))<=>(v1_funct_1(k9_tmap_1(X1,X2))&(v1_funct_2(k9_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))&(v5_pre_topc(k9_tmap_1(X1,X2),X1,k8_tmap_1(X1,X2))&m1_subset_1(k9_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))))))))), inference(assume_negation,[status(cth)],[t122_tmap_1])).
fof(c_0_18, plain, ![X9, X10, X11, X12]:((X11!=k8_tmap_1(X9,X10)|(~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X9)))|(X12!=u1_struct_0(X10)|X11=k6_tmap_1(X9,X12)))|(~v1_pre_topc(X11)|(~v2_pre_topc(X11)|~l1_pre_topc(X11)))|~m1_pre_topc(X10,X9)|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))))&((m1_subset_1(esk3_3(X9,X10,X11),k1_zfmisc_1(u1_struct_0(X9)))|X11=k8_tmap_1(X9,X10)|(~v1_pre_topc(X11)|(~v2_pre_topc(X11)|~l1_pre_topc(X11)))|~m1_pre_topc(X10,X9)|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))))&((esk3_3(X9,X10,X11)=u1_struct_0(X10)|X11=k8_tmap_1(X9,X10)|(~v1_pre_topc(X11)|(~v2_pre_topc(X11)|~l1_pre_topc(X11)))|~m1_pre_topc(X10,X9)|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))))&(X11!=k6_tmap_1(X9,esk3_3(X9,X10,X11))|X11=k8_tmap_1(X9,X10)|(~v1_pre_topc(X11)|(~v2_pre_topc(X11)|~l1_pre_topc(X11)))|~m1_pre_topc(X10,X9)|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d11_tmap_1])])])])])])).
fof(c_0_19, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_pre_topc(esk2_0,esk1_0)&((~v1_tsep_1(esk2_0,esk1_0)|~m1_pre_topc(esk2_0,esk1_0)|(~v1_funct_1(k9_tmap_1(esk1_0,esk2_0))|(~v1_funct_2(k9_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))|(~v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))|~m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))))))&(((v1_funct_1(k9_tmap_1(esk1_0,esk2_0))|v1_tsep_1(esk2_0,esk1_0))&((v1_funct_2(k9_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))|v1_tsep_1(esk2_0,esk1_0))&((v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))|v1_tsep_1(esk2_0,esk1_0))&(m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))|v1_tsep_1(esk2_0,esk1_0)))))&((v1_funct_1(k9_tmap_1(esk1_0,esk2_0))|m1_pre_topc(esk2_0,esk1_0))&((v1_funct_2(k9_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))|m1_pre_topc(esk2_0,esk1_0))&((v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))|m1_pre_topc(esk2_0,esk1_0))&(m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))|m1_pre_topc(esk2_0,esk1_0))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])])).
fof(c_0_20, plain, ![X14, X15]:((~v2_struct_0(k6_tmap_1(X14,X15))|(v2_struct_0(X14)|(~v2_pre_topc(X14)|~l1_pre_topc(X14))|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X14)))))&((v1_pre_topc(k6_tmap_1(X14,X15))|(v2_struct_0(X14)|(~v2_pre_topc(X14)|~l1_pre_topc(X14))|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X14)))))&(v2_pre_topc(k6_tmap_1(X14,X15))|(v2_struct_0(X14)|(~v2_pre_topc(X14)|~l1_pre_topc(X14))|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X14))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc4_tmap_1])])])])).
fof(c_0_21, plain, ![X16, X17]:(~l1_pre_topc(X16)|(~m1_pre_topc(X17,X16)|m1_subset_1(u1_struct_0(X17),k1_zfmisc_1(u1_struct_0(X16))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])).
fof(c_0_22, plain, ![X18, X19]:((v1_pre_topc(k6_tmap_1(X18,X19))|(v2_struct_0(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18))|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))&((v2_pre_topc(k6_tmap_1(X18,X19))|(v2_struct_0(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18))|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))&(l1_pre_topc(k6_tmap_1(X18,X19))|(v2_struct_0(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18))|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_tmap_1])])])])).
cnf(c_0_23, plain, (esk3_3(X1,X2,X3)=u1_struct_0(X2)|X3=k8_tmap_1(X1,X2)|v2_struct_0(X1)|~v1_pre_topc(X3)|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, plain, (v1_pre_topc(k6_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, plain, (l1_pre_topc(k6_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, plain, (v2_pre_topc(k6_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_32, plain, ![X20, X21]:((u1_struct_0(k6_tmap_1(X20,X21))=u1_struct_0(X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20))))&(u1_pre_topc(k6_tmap_1(X20,X21))=k5_tmap_1(X20,X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t104_tmap_1])])])])])).
cnf(c_0_33, plain, (X1=k8_tmap_1(X2,X3)|v2_struct_0(X2)|X1!=k6_tmap_1(X2,esk3_3(X2,X3,X1))|~v1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_34, negated_conjecture, (esk3_3(esk1_0,esk2_0,X1)=u1_struct_0(esk2_0)|X1=k8_tmap_1(esk1_0,esk2_0)|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_35, plain, (v1_pre_topc(k6_tmap_1(X1,u1_struct_0(X2)))|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_36, plain, (l1_pre_topc(k6_tmap_1(X1,u1_struct_0(X2)))|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_30, c_0_29])).
cnf(c_0_37, plain, (v2_pre_topc(k6_tmap_1(X1,u1_struct_0(X2)))|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_31, c_0_29])).
cnf(c_0_38, plain, (u1_struct_0(k6_tmap_1(X1,X2))=u1_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39, negated_conjecture, (X1=k8_tmap_1(esk1_0,esk2_0)|k6_tmap_1(esk1_0,u1_struct_0(esk2_0))!=X1|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_40, negated_conjecture, (v1_pre_topc(k6_tmap_1(esk1_0,u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_41, negated_conjecture, (l1_pre_topc(k6_tmap_1(esk1_0,u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_42, negated_conjecture, (v2_pre_topc(k6_tmap_1(esk1_0,u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_43, plain, (u1_struct_0(k6_tmap_1(X1,u1_struct_0(X2)))=u1_struct_0(X1)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_38, c_0_29])).
cnf(c_0_44, negated_conjecture, (k6_tmap_1(esk1_0,u1_struct_0(esk2_0))=k8_tmap_1(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41]), c_0_42])])).
cnf(c_0_45, negated_conjecture, (u1_struct_0(k8_tmap_1(esk1_0,esk2_0))=u1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_24]), c_0_44]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_46, negated_conjecture, (l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_41, c_0_44])).
fof(c_0_47, plain, ![X22]:(~l1_pre_topc(X22)|m1_pre_topc(X22,X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tsep_1])])).
cnf(c_0_48, negated_conjecture, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X1,k8_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_45]), c_0_46])])).
cnf(c_0_49, plain, (m1_pre_topc(X1,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_50, plain, ![X23, X24, X25]:(~l1_pre_topc(X23)|(~m1_pre_topc(X24,X23)|(~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X23)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t39_pre_topc])])])).
cnf(c_0_51, negated_conjecture, (m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_45]), c_0_46])])).
cnf(c_0_52, plain, (m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_53, negated_conjecture, (u1_struct_0(k6_tmap_1(esk1_0,u1_struct_0(esk1_0)))=u1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_51]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(k6_tmap_1(esk1_0,u1_struct_0(esk1_0)),X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_55, negated_conjecture, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(k6_tmap_1(esk1_0,u1_struct_0(esk1_0)),X2)|~m1_pre_topc(X1,esk1_0)|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_29]), c_0_25])])).
cnf(c_0_56, plain, (v2_struct_0(X1)|~v2_struct_0(k6_tmap_1(X1,X2))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_57, negated_conjecture, (m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_pre_topc(k6_tmap_1(esk1_0,u1_struct_0(esk1_0)),X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_55, c_0_24])).
cnf(c_0_58, negated_conjecture, (l1_pre_topc(k6_tmap_1(esk1_0,u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_51]), c_0_25]), c_0_26])]), c_0_27])).
fof(c_0_59, plain, ![X34, X35]:(v2_struct_0(X34)|(~v2_pre_topc(X34)|~l1_pre_topc(X34))|(~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X34)))|k7_tmap_1(X34,X35)=k6_partfun1(u1_struct_0(X34)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_tmap_1])])])])).
fof(c_0_60, plain, ![X36]:(v2_struct_0(X36)|~l1_struct_0(X36)|~v1_xboole_0(u1_struct_0(X36))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_61, plain, ![X37]:(~l1_pre_topc(X37)|l1_struct_0(X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_62, negated_conjecture, (~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_struct_0(k8_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_44]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_63, negated_conjecture, (m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_49]), c_0_53]), c_0_58])])).
cnf(c_0_64, plain, (v2_struct_0(X1)|k7_tmap_1(X1,X2)=k6_partfun1(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_59])).
fof(c_0_65, plain, ![X38, X39]:((v1_funct_1(k9_tmap_1(X38,X39))|(v2_struct_0(X38)|(~v2_pre_topc(X38)|~l1_pre_topc(X38))|~m1_pre_topc(X39,X38)))&((v1_funct_2(k9_tmap_1(X38,X39),u1_struct_0(X38),u1_struct_0(k8_tmap_1(X38,X39)))|(v2_struct_0(X38)|(~v2_pre_topc(X38)|~l1_pre_topc(X38))|~m1_pre_topc(X39,X38)))&(m1_subset_1(k9_tmap_1(X38,X39),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X38),u1_struct_0(k8_tmap_1(X38,X39)))))|(v2_struct_0(X38)|(~v2_pre_topc(X38)|~l1_pre_topc(X38))|~m1_pre_topc(X39,X38))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k9_tmap_1])])])])).
cnf(c_0_66, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_67, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_68, plain, (u1_struct_0(k6_tmap_1(X1,u1_struct_0(X1)))=u1_struct_0(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_43, c_0_49])).
cnf(c_0_69, negated_conjecture, (v2_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_42, c_0_44])).
cnf(c_0_70, negated_conjecture, (~v2_struct_0(k8_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_63])])).
cnf(c_0_71, plain, (l1_pre_topc(k6_tmap_1(X1,u1_struct_0(X1)))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_36, c_0_49])).
fof(c_0_72, plain, ![X40, X41, X42, X43]:((X42!=k9_tmap_1(X40,X41)|(~m1_subset_1(X43,k1_zfmisc_1(u1_struct_0(X40)))|(X43!=u1_struct_0(X41)|r1_funct_2(u1_struct_0(X40),u1_struct_0(k8_tmap_1(X40,X41)),u1_struct_0(X40),u1_struct_0(k6_tmap_1(X40,X43)),X42,k7_tmap_1(X40,X43))))|(~v1_funct_1(X42)|(~v1_funct_2(X42,u1_struct_0(X40),u1_struct_0(k8_tmap_1(X40,X41)))|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X40),u1_struct_0(k8_tmap_1(X40,X41)))))))|~m1_pre_topc(X41,X40)|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40))))&((m1_subset_1(esk5_3(X40,X41,X42),k1_zfmisc_1(u1_struct_0(X40)))|X42=k9_tmap_1(X40,X41)|(~v1_funct_1(X42)|(~v1_funct_2(X42,u1_struct_0(X40),u1_struct_0(k8_tmap_1(X40,X41)))|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X40),u1_struct_0(k8_tmap_1(X40,X41)))))))|~m1_pre_topc(X41,X40)|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40))))&((esk5_3(X40,X41,X42)=u1_struct_0(X41)|X42=k9_tmap_1(X40,X41)|(~v1_funct_1(X42)|(~v1_funct_2(X42,u1_struct_0(X40),u1_struct_0(k8_tmap_1(X40,X41)))|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X40),u1_struct_0(k8_tmap_1(X40,X41)))))))|~m1_pre_topc(X41,X40)|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40))))&(~r1_funct_2(u1_struct_0(X40),u1_struct_0(k8_tmap_1(X40,X41)),u1_struct_0(X40),u1_struct_0(k6_tmap_1(X40,esk5_3(X40,X41,X42))),X42,k7_tmap_1(X40,esk5_3(X40,X41,X42)))|X42=k9_tmap_1(X40,X41)|(~v1_funct_1(X42)|(~v1_funct_2(X42,u1_struct_0(X40),u1_struct_0(k8_tmap_1(X40,X41)))|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X40),u1_struct_0(k8_tmap_1(X40,X41)))))))|~m1_pre_topc(X41,X40)|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d12_tmap_1])])])])])])).
fof(c_0_73, plain, ![X32, X33]:((v1_funct_1(k7_tmap_1(X32,X33))|(v2_struct_0(X32)|(~v2_pre_topc(X32)|~l1_pre_topc(X32))|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32)))))&((v1_funct_2(k7_tmap_1(X32,X33),u1_struct_0(X32),u1_struct_0(k6_tmap_1(X32,X33)))|(v2_struct_0(X32)|(~v2_pre_topc(X32)|~l1_pre_topc(X32))|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32)))))&(m1_subset_1(k7_tmap_1(X32,X33),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X32),u1_struct_0(k6_tmap_1(X32,X33)))))|(v2_struct_0(X32)|(~v2_pre_topc(X32)|~l1_pre_topc(X32))|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_tmap_1])])])])).
cnf(c_0_74, negated_conjecture, (k6_partfun1(u1_struct_0(esk1_0))=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_51]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_75, plain, (m1_subset_1(k9_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_76, plain, (v1_funct_2(k9_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_77, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_78, negated_conjecture, (u1_struct_0(k6_tmap_1(k8_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0)))=u1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_45]), c_0_46]), c_0_69])]), c_0_70])).
cnf(c_0_79, negated_conjecture, (l1_pre_topc(k6_tmap_1(k8_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_45]), c_0_46]), c_0_69])]), c_0_70])).
cnf(c_0_80, plain, (esk5_3(X1,X2,X3)=u1_struct_0(X2)|X3=k9_tmap_1(X1,X2)|v2_struct_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_81, plain, (v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_82, plain, (m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_83, negated_conjecture, (k7_tmap_1(esk1_0,u1_struct_0(esk2_0))=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_63]), c_0_74]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_84, plain, (v1_funct_1(k7_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_73])).
fof(c_0_85, plain, ![X45, X46, X47, X48, X49, X50]:(v1_xboole_0(X46)|(v1_xboole_0(X48)|(~v1_funct_1(X49)|(~v1_funct_2(X49,X45,X46)|~m1_subset_1(X49,k1_zfmisc_1(k2_zfmisc_1(X45,X46))))|(~v1_funct_1(X50)|(~v1_funct_2(X50,X47,X48)|~m1_subset_1(X50,k1_zfmisc_1(k2_zfmisc_1(X47,X48)))))))|r1_funct_2(X45,X46,X47,X48,X49,X49)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r1_funct_2])])])).
cnf(c_0_86, negated_conjecture, (m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_87, negated_conjecture, (v1_funct_2(k9_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_88, plain, (v1_funct_1(k9_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_89, negated_conjecture, (v2_struct_0(k6_tmap_1(k8_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0)))|~v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_79])])).
fof(c_0_90, plain, ![X30, X31]:(((v1_funct_1(k7_tmap_1(X30,X31))|~v3_pre_topc(X31,X30)|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))|(v2_struct_0(X30)|(~v2_pre_topc(X30)|~l1_pre_topc(X30))))&((v1_funct_2(k7_tmap_1(X30,X31),u1_struct_0(X30),u1_struct_0(k6_tmap_1(X30,X31)))|~v3_pre_topc(X31,X30)|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))|(v2_struct_0(X30)|(~v2_pre_topc(X30)|~l1_pre_topc(X30))))&((v5_pre_topc(k7_tmap_1(X30,X31),X30,k6_tmap_1(X30,X31))|~v3_pre_topc(X31,X30)|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))|(v2_struct_0(X30)|(~v2_pre_topc(X30)|~l1_pre_topc(X30))))&(m1_subset_1(k7_tmap_1(X30,X31),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X30),u1_struct_0(k6_tmap_1(X30,X31)))))|~v3_pre_topc(X31,X30)|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))|(v2_struct_0(X30)|(~v2_pre_topc(X30)|~l1_pre_topc(X30)))))))&(~v1_funct_1(k7_tmap_1(X30,X31))|(~v1_funct_2(k7_tmap_1(X30,X31),u1_struct_0(X30),u1_struct_0(k6_tmap_1(X30,X31)))|(~v5_pre_topc(k7_tmap_1(X30,X31),X30,k6_tmap_1(X30,X31))|~m1_subset_1(k7_tmap_1(X30,X31),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X30),u1_struct_0(k6_tmap_1(X30,X31)))))))|v3_pre_topc(X31,X30)|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))|(v2_struct_0(X30)|(~v2_pre_topc(X30)|~l1_pre_topc(X30))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t113_tmap_1])])])])])).
cnf(c_0_91, plain, (X3=k9_tmap_1(X1,X2)|v2_struct_0(X1)|~r1_funct_2(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,esk5_3(X1,X2,X3))),X3,k7_tmap_1(X1,esk5_3(X1,X2,X3)))|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_92, negated_conjecture, (esk5_3(esk1_0,esk2_0,X1)=u1_struct_0(esk2_0)|X1=k9_tmap_1(esk1_0,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_45]), c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_93, negated_conjecture, (v1_funct_2(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_51]), c_0_25]), c_0_26])]), c_0_27]), c_0_53])).
cnf(c_0_94, negated_conjecture, (m1_subset_1(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_63]), c_0_44]), c_0_45]), c_0_25]), c_0_26])]), c_0_27]), c_0_83])).
cnf(c_0_95, negated_conjecture, (v1_funct_1(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_51]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_96, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|r1_funct_2(X4,X1,X6,X2,X3,X3)|~v1_funct_1(X3)|~v1_funct_2(X3,X4,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X1)))|~v1_funct_1(X5)|~v1_funct_2(X5,X6,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X2)))), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_97, negated_conjecture, (m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(rw,[status(thm)],[c_0_86, c_0_45])).
cnf(c_0_98, negated_conjecture, (v1_funct_2(k9_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(rw,[status(thm)],[c_0_87, c_0_45])).
cnf(c_0_99, negated_conjecture, (v1_funct_1(k9_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_100, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_89]), c_0_45]), c_0_51]), c_0_46]), c_0_69])]), c_0_70])).
cnf(c_0_101, negated_conjecture, (~v1_tsep_1(esk2_0,esk1_0)|~m1_pre_topc(esk2_0,esk1_0)|~v1_funct_1(k9_tmap_1(esk1_0,esk2_0))|~v1_funct_2(k9_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))|~v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))|~m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_102, plain, (v3_pre_topc(X2,X1)|v2_struct_0(X1)|~v1_funct_1(k7_tmap_1(X1,X2))|~v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))|~v5_pre_topc(k7_tmap_1(X1,X2),X1,k6_tmap_1(X1,X2))|~m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_103, negated_conjecture, (X1=k9_tmap_1(esk1_0,esk2_0)|~r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk5_3(esk1_0,esk2_0,X1))),X1,k7_tmap_1(esk1_0,esk5_3(esk1_0,esk2_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_45]), c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_104, negated_conjecture, (esk5_3(esk1_0,esk2_0,k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))=u1_struct_0(esk2_0)|k7_tmap_1(esk1_0,u1_struct_0(esk1_0))=k9_tmap_1(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_94]), c_0_95])])).
cnf(c_0_105, negated_conjecture, (r1_funct_2(X1,X2,u1_struct_0(esk1_0),u1_struct_0(esk1_0),X3,X3)|v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_98]), c_0_99])]), c_0_100])).
fof(c_0_106, plain, ![X26, X27, X28]:((~v1_tsep_1(X27,X26)|(~m1_subset_1(X28,k1_zfmisc_1(u1_struct_0(X26)))|(X28!=u1_struct_0(X27)|v3_pre_topc(X28,X26)))|~m1_pre_topc(X27,X26)|~l1_pre_topc(X26))&((m1_subset_1(esk4_2(X26,X27),k1_zfmisc_1(u1_struct_0(X26)))|v1_tsep_1(X27,X26)|~m1_pre_topc(X27,X26)|~l1_pre_topc(X26))&((esk4_2(X26,X27)=u1_struct_0(X27)|v1_tsep_1(X27,X26)|~m1_pre_topc(X27,X26)|~l1_pre_topc(X26))&(~v3_pre_topc(esk4_2(X26,X27),X26)|v1_tsep_1(X27,X26)|~m1_pre_topc(X27,X26)|~l1_pre_topc(X26))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tsep_1])])])])])).
cnf(c_0_107, negated_conjecture, (~m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))|~v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))|~v1_funct_2(k9_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))|~v1_funct_1(k9_tmap_1(esk1_0,esk2_0))|~v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_101, c_0_24])])).
cnf(c_0_108, plain, (v3_pre_topc(X1,X2)|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v5_pre_topc(k7_tmap_1(X2,X1),X2,k6_tmap_1(X2,X1))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_102, c_0_84]), c_0_81]), c_0_82])).
cnf(c_0_109, negated_conjecture, (k7_tmap_1(esk1_0,u1_struct_0(esk1_0))=k9_tmap_1(esk1_0,esk2_0)|~r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_44]), c_0_45]), c_0_83]), c_0_94]), c_0_93]), c_0_95])])).
cnf(c_0_110, negated_conjecture, (r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_94]), c_0_93]), c_0_95])]), c_0_100])).
cnf(c_0_111, plain, (esk4_2(X1,X2)=u1_struct_0(X2)|v1_tsep_1(X2,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_112, negated_conjecture, (~m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))|~v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))|~v1_funct_2(k9_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))|~v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_107, c_0_99])])).
cnf(c_0_113, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),esk1_0)|~v5_pre_topc(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),esk1_0,k8_tmap_1(esk1_0,esk2_0))), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_63]), c_0_44]), c_0_25]), c_0_26])]), c_0_27]), c_0_83])).
cnf(c_0_114, negated_conjecture, (k7_tmap_1(esk1_0,u1_struct_0(esk1_0))=k9_tmap_1(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_109, c_0_110])])).
cnf(c_0_115, plain, (v1_tsep_1(X2,X1)|~v3_pre_topc(esk4_2(X1,X2),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_116, negated_conjecture, (esk4_2(esk1_0,esk2_0)=u1_struct_0(esk2_0)|v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_24]), c_0_25])])).
cnf(c_0_117, negated_conjecture, (~m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))|~v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))|~v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_112, c_0_87])])).
cnf(c_0_118, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),esk1_0)|~v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_113, c_0_114])).
cnf(c_0_119, negated_conjecture, (v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))|v1_tsep_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_120, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|~v3_pre_topc(u1_struct_0(esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_24]), c_0_25])])).
cnf(c_0_121, plain, (v5_pre_topc(k7_tmap_1(X1,X2),X1,k6_tmap_1(X1,X2))|v2_struct_0(X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_122, negated_conjecture, (~v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))|~v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_117, c_0_86])])).
cnf(c_0_123, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_120])).
cnf(c_0_124, negated_conjecture, (v5_pre_topc(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),esk1_0,k8_tmap_1(esk1_0,esk2_0))|~v3_pre_topc(u1_struct_0(esk2_0),esk1_0)), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_63]), c_0_44]), c_0_25]), c_0_26])]), c_0_27]), c_0_83])).
cnf(c_0_125, negated_conjecture, (~v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_122, c_0_123])])).
cnf(c_0_126, plain, (v3_pre_topc(X3,X2)|~v1_tsep_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|X3!=u1_struct_0(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_127, negated_conjecture, (~v3_pre_topc(u1_struct_0(esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_124, c_0_114]), c_0_125])).
cnf(c_0_128, plain, (v3_pre_topc(u1_struct_0(X1),X2)|~v1_tsep_1(X1,X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_126]), c_0_29])).
cnf(c_0_129, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128]), c_0_123]), c_0_24]), c_0_25])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 130
# Proof object clause steps            : 95
# Proof object formula steps           : 35
# Proof object conjectures             : 62
# Proof object clause conjectures      : 59
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 33
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 48
# Proof object simplifying inferences  : 179
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 51
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 51
# Processed clauses                    : 1906
# ...of these trivial                  : 27
# ...subsumed                          : 941
# ...remaining for further processing  : 938
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 120
# Backward-rewritten                   : 299
# Generated clauses                    : 5162
# ...of the previous two non-trivial   : 4359
# Contextual simplify-reflections      : 169
# Paramodulations                      : 5144
# Factorizations                       : 0
# Equation resolutions                 : 11
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 470
#    Positive orientable unit clauses  : 32
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 433
# Current number of unprocessed clauses: 2197
# ...number of literals in the above   : 12534
# Current number of archived formulas  : 0
# Current number of archived clauses   : 468
# Clause-clause subsumption calls (NU) : 110130
# Rec. Clause-clause subsumption calls : 24920
# Non-unit clause-clause subsumptions  : 1199
# Unit Clause-clause subsumption calls : 1350
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 70
# BW rewrite match successes           : 33
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 175034

# -------------------------------------------------
# User time                : 0.127 s
# System time              : 0.010 s
# Total time               : 0.137 s
# Maximum resident set size: 28816 pages
