# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/11b6a5d7b31c491f8b8163b2cb2406e8.E_input', d2_tarski)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/11b6a5d7b31c491f8b8163b2cb2406e8.E_input', t70_enumset1)).
fof(t51_mcart_1, conjecture, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(~(X4=k5_mcart_1(X1,X2,X3,X4))&(~(X4=k6_mcart_1(X1,X2,X3,X4))&~(X4=k7_mcart_1(X1,X2,X3,X4)))))))))), file('data/MPTP/proofs/11b6a5d7b31c491f8b8163b2cb2406e8.E_input', t51_mcart_1)).
fof(t9_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k4_tarski(X3,X4))))))), file('data/MPTP/proofs/11b6a5d7b31c491f8b8163b2cb2406e8.E_input', t9_mcart_1)).
fof(t48_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>X4=k3_mcart_1(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4),k7_mcart_1(X1,X2,X3,X4)))))))), file('data/MPTP/proofs/11b6a5d7b31c491f8b8163b2cb2406e8.E_input', t48_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/11b6a5d7b31c491f8b8163b2cb2406e8.E_input', d3_zfmisc_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/11b6a5d7b31c491f8b8163b2cb2406e8.E_input', d3_mcart_1)).
fof(t20_mcart_1, axiom, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1)))), file('data/MPTP/proofs/11b6a5d7b31c491f8b8163b2cb2406e8.E_input', t20_mcart_1)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/11b6a5d7b31c491f8b8163b2cb2406e8.E_input', t7_mcart_1)).
fof(t50_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))&(k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))&k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4))))))))), file('data/MPTP/proofs/11b6a5d7b31c491f8b8163b2cb2406e8.E_input', t50_mcart_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/11b6a5d7b31c491f8b8163b2cb2406e8.E_input', t113_zfmisc_1)).
fof(t152_zfmisc_1, axiom, ![X1, X2]:~(r2_hidden(X1,k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/11b6a5d7b31c491f8b8163b2cb2406e8.E_input', t152_zfmisc_1)).
fof(c_0_12, plain, ![X5, X6, X7, X8, X9, X10, X11, X12]:(((~r2_hidden(X8,X7)|(X8=X5|X8=X6)|X7!=k2_tarski(X5,X6))&((X9!=X5|r2_hidden(X9,X7)|X7!=k2_tarski(X5,X6))&(X9!=X6|r2_hidden(X9,X7)|X7!=k2_tarski(X5,X6))))&(((esk1_3(X10,X11,X12)!=X10|~r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k2_tarski(X10,X11))&(esk1_3(X10,X11,X12)!=X11|~r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k2_tarski(X10,X11)))&(r2_hidden(esk1_3(X10,X11,X12),X12)|(esk1_3(X10,X11,X12)=X10|esk1_3(X10,X11,X12)=X11)|X12=k2_tarski(X10,X11)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_13, plain, ![X14, X15]:k1_enumset1(X14,X14,X15)=k2_tarski(X14,X15), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_14, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_16, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_17, negated_conjecture, ~(![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(~(X4=k5_mcart_1(X1,X2,X3,X4))&(~(X4=k6_mcart_1(X1,X2,X3,X4))&~(X4=k7_mcart_1(X1,X2,X3,X4))))))))))), inference(assume_negation,[status(cth)],[t51_mcart_1])).
cnf(c_0_18, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_19, plain, ![X20, X22, X23]:((r2_hidden(esk6_1(X20),X20)|X20=k1_xboole_0)&((~r2_hidden(X22,X20)|esk6_1(X20)!=k4_tarski(X22,X23)|X20=k1_xboole_0)&(~r2_hidden(X23,X20)|esk6_1(X20)!=k4_tarski(X22,X23)|X20=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_mcart_1])])])])])])).
cnf(c_0_20, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X2,X2,X4)), inference(rw,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_21, plain, (X1=X4|X1=X3|X2!=k1_enumset1(X3,X3,X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_16, c_0_15])).
fof(c_0_22, plain, ![X24, X25, X26, X27]:(X24=k1_xboole_0|(X25=k1_xboole_0|(X26=k1_xboole_0|(~m1_subset_1(X27,k3_zfmisc_1(X24,X25,X26))|X27=k3_mcart_1(k5_mcart_1(X24,X25,X26,X27),k6_mcart_1(X24,X25,X26,X27),k7_mcart_1(X24,X25,X26,X27)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t48_mcart_1])])])])).
fof(c_0_23, plain, ![X28, X29, X30]:k3_zfmisc_1(X28,X29,X30)=k2_zfmisc_1(k2_zfmisc_1(X28,X29),X30), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_24, plain, ![X31, X32, X33]:k3_mcart_1(X31,X32,X33)=k4_tarski(k4_tarski(X31,X32),X33), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_25, negated_conjecture, (esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&(esk4_0!=k1_xboole_0&(m1_subset_1(esk5_0,k3_zfmisc_1(esk2_0,esk3_0,esk4_0))&(esk5_0=k5_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)|(esk5_0=k6_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)|esk5_0=k7_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])).
fof(c_0_26, plain, ![X34, X35, X36]:((X34!=k1_mcart_1(X34)|X34!=k4_tarski(X35,X36))&(X34!=k2_mcart_1(X34)|X34!=k4_tarski(X35,X36))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t20_mcart_1])])])])])).
fof(c_0_27, plain, ![X37, X38]:(k1_mcart_1(k4_tarski(X37,X38))=X37&k2_mcart_1(k4_tarski(X37,X38))=X38), inference(variable_rename,[status(thm)],[t7_mcart_1])).
fof(c_0_28, plain, ![X39, X40, X41, X42]:((k5_mcart_1(X39,X40,X41,X42)=k1_mcart_1(k1_mcart_1(X42))|~m1_subset_1(X42,k3_zfmisc_1(X39,X40,X41))|X41=k1_xboole_0|X40=k1_xboole_0|X39=k1_xboole_0)&((k6_mcart_1(X39,X40,X41,X42)=k2_mcart_1(k1_mcart_1(X42))|~m1_subset_1(X42,k3_zfmisc_1(X39,X40,X41))|X41=k1_xboole_0|X40=k1_xboole_0|X39=k1_xboole_0)&(k7_mcart_1(X39,X40,X41,X42)=k2_mcart_1(X42)|~m1_subset_1(X42,k3_zfmisc_1(X39,X40,X41))|X41=k1_xboole_0|X40=k1_xboole_0|X39=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_mcart_1])])])])])).
cnf(c_0_29, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X4,X2)), inference(rw,[status(thm)],[c_0_18, c_0_15])).
cnf(c_0_30, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk6_1(X2)!=k4_tarski(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_31, plain, (r2_hidden(X1,k1_enumset1(X1,X1,X2))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_20])])).
cnf(c_0_32, plain, (X1=X2|X1=X3|~r2_hidden(X1,k1_enumset1(X3,X3,X2))), inference(er,[status(thm)],[c_0_21])).
cnf(c_0_33, plain, (r2_hidden(esk6_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_34, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k3_mcart_1(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4),k7_mcart_1(X1,X2,X3,X4))|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_35, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_36, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_37, negated_conjecture, (m1_subset_1(esk5_0,k3_zfmisc_1(esk2_0,esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_38, plain, (X1!=k2_mcart_1(X1)|X1!=k4_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_39, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_40, plain, (k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_41, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk6_1(X2)!=k4_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_42, plain, (r2_hidden(X1,k1_enumset1(X2,X2,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_29])])).
cnf(c_0_43, plain, (k1_enumset1(X1,X1,X2)=k1_xboole_0|esk6_1(k1_enumset1(X1,X1,X2))!=k4_tarski(X3,X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_44, plain, (esk6_1(k1_enumset1(X1,X1,X2))=X2|esk6_1(k1_enumset1(X1,X1,X2))=X1|k1_enumset1(X1,X1,X2)=k1_xboole_0), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_45, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X4=k4_tarski(k4_tarski(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4)),k7_mcart_1(X1,X2,X3,X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_35]), c_0_36])).
cnf(c_0_46, negated_conjecture, (esk5_0=k5_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)|esk5_0=k6_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)|esk5_0=k7_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_47, negated_conjecture, (m1_subset_1(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0))), inference(rw,[status(thm)],[c_0_37, c_0_35])).
cnf(c_0_48, plain, (k4_tarski(X1,X2)!=X2), inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_38]), c_0_39])).
cnf(c_0_49, negated_conjecture, (esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_50, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_51, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_52, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_40, c_0_35])).
fof(c_0_53, plain, ![X43, X44]:((k2_zfmisc_1(X43,X44)!=k1_xboole_0|(X43=k1_xboole_0|X44=k1_xboole_0))&((X43!=k1_xboole_0|k2_zfmisc_1(X43,X44)=k1_xboole_0)&(X44!=k1_xboole_0|k2_zfmisc_1(X43,X44)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_54, plain, (k1_enumset1(X1,X1,X2)=k1_xboole_0|esk6_1(k1_enumset1(X1,X1,X2))!=k4_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_55, plain, (esk6_1(k1_enumset1(X1,X1,k4_tarski(X2,X1)))=X1|k1_enumset1(X1,X1,k4_tarski(X2,X1))=k1_xboole_0), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44])])).
cnf(c_0_56, negated_conjecture, (k6_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)=esk5_0|k5_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)=esk5_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])]), c_0_48]), c_0_49]), c_0_50]), c_0_51])).
cnf(c_0_57, negated_conjecture, (k7_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)=k2_mcart_1(esk5_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_47]), c_0_49]), c_0_50]), c_0_51])).
fof(c_0_58, plain, ![X45, X46]:~r2_hidden(X45,k2_zfmisc_1(X45,X46)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_zfmisc_1])])).
cnf(c_0_59, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_60, plain, (k1_enumset1(X1,X1,k4_tarski(X2,X1))=k1_xboole_0|k4_tarski(k4_tarski(X2,X1),X3)!=X1), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_61, negated_conjecture, (k4_tarski(k4_tarski(k5_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0),esk5_0),k2_mcart_1(esk5_0))=esk5_0|k5_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)=esk5_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_56]), c_0_57]), c_0_47])]), c_0_49]), c_0_50]), c_0_51])).
cnf(c_0_62, plain, (~r2_hidden(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_63, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_59])).
cnf(c_0_64, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_65, negated_conjecture, (k4_tarski(k4_tarski(k5_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0),k6_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)),k2_mcart_1(esk5_0))=esk5_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_57]), c_0_47])]), c_0_49]), c_0_50]), c_0_51])).
cnf(c_0_66, negated_conjecture, (k1_enumset1(esk5_0,esk5_0,k4_tarski(k5_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0),esk5_0))=k1_xboole_0|k5_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)=esk5_0), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_67, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_68, plain, (X1!=k1_mcart_1(X1)|X1!=k4_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_69, negated_conjecture, (k4_tarski(k5_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0),k6_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0))=k1_mcart_1(esk5_0)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_70, plain, (k1_enumset1(X1,X1,X2)=k1_xboole_0|esk6_1(k1_enumset1(X1,X1,X2))!=k4_tarski(X1,X3)), inference(spm,[status(thm)],[c_0_41, c_0_31])).
cnf(c_0_71, negated_conjecture, (k5_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0)=esk5_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_66]), c_0_67])).
cnf(c_0_72, plain, (k4_tarski(X1,X2)!=X1), inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_68]), c_0_64])).
cnf(c_0_73, plain, (esk6_1(k1_enumset1(k4_tarski(X1,X2),k4_tarski(X1,X2),X1))=X1|k1_enumset1(k4_tarski(X1,X2),k4_tarski(X1,X2),X1)=k1_xboole_0), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_44])])).
cnf(c_0_74, negated_conjecture, (k4_tarski(k1_mcart_1(esk5_0),k2_mcart_1(esk5_0))=esk5_0), inference(rw,[status(thm)],[c_0_65, c_0_69])).
cnf(c_0_75, plain, (esk6_1(k1_enumset1(X1,X1,k4_tarski(X1,X2)))=X1|k1_enumset1(X1,X1,k4_tarski(X1,X2))=k1_xboole_0), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_44])])).
cnf(c_0_76, negated_conjecture, (k4_tarski(esk5_0,k6_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0))=k1_mcart_1(esk5_0)), inference(rw,[status(thm)],[c_0_69, c_0_71])).
cnf(c_0_77, negated_conjecture, (k4_tarski(k5_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0),k6_mcart_1(esk2_0,esk3_0,esk4_0,esk5_0))!=esk5_0), inference(spm,[status(thm)],[c_0_72, c_0_65])).
cnf(c_0_78, negated_conjecture, (esk6_1(k1_enumset1(esk5_0,esk5_0,k1_mcart_1(esk5_0)))=k1_mcart_1(esk5_0)|k1_enumset1(esk5_0,esk5_0,k1_mcart_1(esk5_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_79, negated_conjecture, (esk6_1(k1_enumset1(esk5_0,esk5_0,k1_mcart_1(esk5_0)))=esk5_0|k1_enumset1(esk5_0,esk5_0,k1_mcart_1(esk5_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_80, negated_conjecture, (k1_mcart_1(esk5_0)!=esk5_0), inference(rw,[status(thm)],[c_0_77, c_0_69])).
cnf(c_0_81, negated_conjecture, (k1_enumset1(esk5_0,esk5_0,k1_mcart_1(esk5_0))=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80])).
cnf(c_0_82, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_81]), c_0_67]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 83
# Proof object clause steps            : 58
# Proof object formula steps           : 25
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 46
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 26
# Processed clauses                    : 155
# ...of these trivial                  : 3
# ...subsumed                          : 43
# ...remaining for further processing  : 109
# Other redundant clauses eliminated   : 47
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 11
# Generated clauses                    : 460
# ...of the previous two non-trivial   : 386
# Contextual simplify-reflections      : 0
# Paramodulations                      : 399
# Factorizations                       : 16
# Equation resolutions                 : 47
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 62
#    Positive orientable unit clauses  : 15
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 11
#    Non-unit-clauses                  : 36
# Current number of unprocessed clauses: 269
# ...number of literals in the above   : 1300
# Current number of archived formulas  : 0
# Current number of archived clauses   : 43
# Clause-clause subsumption calls (NU) : 167
# Rec. Clause-clause subsumption calls : 122
# Non-unit clause-clause subsumptions  : 12
# Unit Clause-clause subsumption calls : 73
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7973

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.004 s
# Total time               : 0.015 s
# Maximum resident set size: 28684 pages
