# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t113_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>(v1_funct_1(k7_tmap_1(X1,X2))&(v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))&(v5_pre_topc(k7_tmap_1(X1,X2),X1,k6_tmap_1(X1,X2))&m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2))))))))))), file('data/MPTP/proofs/11d136ec1928407181477f53fc79fad6.E_input', t113_tmap_1)).
fof(dt_k7_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_funct_1(k7_tmap_1(X1,X2))&(v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))&m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))))), file('data/MPTP/proofs/11d136ec1928407181477f53fc79fad6.E_input', dt_k7_tmap_1)).
fof(t104_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(u1_struct_0(k6_tmap_1(X1,X2))=u1_struct_0(X1)&u1_pre_topc(k6_tmap_1(X1,X2))=k5_tmap_1(X1,X2)))), file('data/MPTP/proofs/11d136ec1928407181477f53fc79fad6.E_input', t104_tmap_1)).
fof(d10_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k7_tmap_1(X1,X2)=k6_partfun1(u1_struct_0(X1)))), file('data/MPTP/proofs/11d136ec1928407181477f53fc79fad6.E_input', d10_tmap_1)).
fof(dt_k6_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_pre_topc(k6_tmap_1(X1,X2))&(v2_pre_topc(k6_tmap_1(X1,X2))&l1_pre_topc(k6_tmap_1(X1,X2))))), file('data/MPTP/proofs/11d136ec1928407181477f53fc79fad6.E_input', dt_k6_tmap_1)).
fof(t105_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k6_tmap_1(X1,X2))))=>(X3=X2=>v3_pre_topc(X3,k6_tmap_1(X1,X2)))))), file('data/MPTP/proofs/11d136ec1928407181477f53fc79fad6.E_input', t105_tmap_1)).
fof(t63_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((v2_pre_topc(X2)&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))))))=>(X3=X4=>(v5_pre_topc(X3,X1,X2)<=>v5_pre_topc(X4,X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))))))), file('data/MPTP/proofs/11d136ec1928407181477f53fc79fad6.E_input', t63_pre_topc)).
fof(t106_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=k6_tmap_1(X1,X2)))), file('data/MPTP/proofs/11d136ec1928407181477f53fc79fad6.E_input', t106_tmap_1)).
fof(t55_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>((k2_struct_0(X2)=k1_xboole_0=>k2_struct_0(X1)=k1_xboole_0)=>(v5_pre_topc(X3,X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(v3_pre_topc(X4,X2)=>v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X1)))))))), file('data/MPTP/proofs/11d136ec1928407181477f53fc79fad6.E_input', t55_tops_2)).
fof(t171_funct_2, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k8_relset_1(X1,X1,k6_partfun1(X1),X2)=X2), file('data/MPTP/proofs/11d136ec1928407181477f53fc79fad6.E_input', t171_funct_2)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/11d136ec1928407181477f53fc79fad6.E_input', d3_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/11d136ec1928407181477f53fc79fad6.E_input', dt_l1_pre_topc)).
fof(c_0_12, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>(v1_funct_1(k7_tmap_1(X1,X2))&(v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))&(v5_pre_topc(k7_tmap_1(X1,X2),X1,k6_tmap_1(X1,X2))&m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))))))))), inference(assume_negation,[status(cth)],[t113_tmap_1])).
fof(c_0_13, plain, ![X20, X21]:((v1_funct_1(k7_tmap_1(X20,X21))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))))&((v1_funct_2(k7_tmap_1(X20,X21),u1_struct_0(X20),u1_struct_0(k6_tmap_1(X20,X21)))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))))&(m1_subset_1(k7_tmap_1(X20,X21),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X20),u1_struct_0(k6_tmap_1(X20,X21)))))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_tmap_1])])])])).
fof(c_0_14, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&((~v3_pre_topc(esk2_0,esk1_0)|(~v1_funct_1(k7_tmap_1(esk1_0,esk2_0))|(~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))|(~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))))))))&((v1_funct_1(k7_tmap_1(esk1_0,esk2_0))|v3_pre_topc(esk2_0,esk1_0))&((v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))|v3_pre_topc(esk2_0,esk1_0))&((v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|v3_pre_topc(esk2_0,esk1_0))&(m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))))|v3_pre_topc(esk2_0,esk1_0)))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])])).
cnf(c_0_15, plain, (v1_funct_1(k7_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_16, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_17, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_20, plain, ![X16, X17]:((u1_struct_0(k6_tmap_1(X16,X17))=u1_struct_0(X16)|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))))&(u1_pre_topc(k6_tmap_1(X16,X17))=k5_tmap_1(X16,X17)|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t104_tmap_1])])])])])).
cnf(c_0_21, negated_conjecture, (~v3_pre_topc(esk2_0,esk1_0)|~v1_funct_1(k7_tmap_1(esk1_0,esk2_0))|~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))|~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, negated_conjecture, (v1_funct_1(k7_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_18])]), c_0_19])).
cnf(c_0_23, plain, (u1_struct_0(k6_tmap_1(X1,X2))=u1_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_24, plain, ![X12, X13]:(v2_struct_0(X12)|(~v2_pre_topc(X12)|~l1_pre_topc(X12))|(~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|k7_tmap_1(X12,X13)=k6_partfun1(u1_struct_0(X12)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_tmap_1])])])])).
fof(c_0_25, plain, ![X7, X8]:((v1_pre_topc(k6_tmap_1(X7,X8))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))))&((v2_pre_topc(k6_tmap_1(X7,X8))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))))&(l1_pre_topc(k6_tmap_1(X7,X8))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_tmap_1])])])])).
fof(c_0_26, plain, ![X9, X10, X11]:(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))|(~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|(~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(k6_tmap_1(X9,X10))))|(X11!=X10|v3_pre_topc(X11,k6_tmap_1(X9,X10)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t105_tmap_1])])])])).
cnf(c_0_27, negated_conjecture, (~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))|~v3_pre_topc(esk2_0,esk1_0)|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_22])])).
cnf(c_0_28, negated_conjecture, (u1_struct_0(k6_tmap_1(esk1_0,esk2_0))=u1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_16]), c_0_17]), c_0_18])]), c_0_19])).
cnf(c_0_29, plain, (v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_30, plain, ![X27, X28, X29, X30]:((~v5_pre_topc(X29,X27,X28)|v5_pre_topc(X30,X27,g1_pre_topc(u1_struct_0(X28),u1_pre_topc(X28)))|X29!=X30|(~v1_funct_1(X30)|(~v1_funct_2(X30,u1_struct_0(X27),u1_struct_0(g1_pre_topc(u1_struct_0(X28),u1_pre_topc(X28))))|~m1_subset_1(X30,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X27),u1_struct_0(g1_pre_topc(u1_struct_0(X28),u1_pre_topc(X28))))))))|(~v1_funct_1(X29)|(~v1_funct_2(X29,u1_struct_0(X27),u1_struct_0(X28))|~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X27),u1_struct_0(X28))))))|(~v2_pre_topc(X28)|~l1_pre_topc(X28))|(~v2_pre_topc(X27)|~l1_pre_topc(X27)))&(~v5_pre_topc(X30,X27,g1_pre_topc(u1_struct_0(X28),u1_pre_topc(X28)))|v5_pre_topc(X29,X27,X28)|X29!=X30|(~v1_funct_1(X30)|(~v1_funct_2(X30,u1_struct_0(X27),u1_struct_0(g1_pre_topc(u1_struct_0(X28),u1_pre_topc(X28))))|~m1_subset_1(X30,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X27),u1_struct_0(g1_pre_topc(u1_struct_0(X28),u1_pre_topc(X28))))))))|(~v1_funct_1(X29)|(~v1_funct_2(X29,u1_struct_0(X27),u1_struct_0(X28))|~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X27),u1_struct_0(X28))))))|(~v2_pre_topc(X28)|~l1_pre_topc(X28))|(~v2_pre_topc(X27)|~l1_pre_topc(X27)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_pre_topc])])])])).
fof(c_0_31, plain, ![X18, X19]:((~v3_pre_topc(X19,X18)|g1_pre_topc(u1_struct_0(X18),u1_pre_topc(X18))=k6_tmap_1(X18,X19)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|(v2_struct_0(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18))))&(g1_pre_topc(u1_struct_0(X18),u1_pre_topc(X18))!=k6_tmap_1(X18,X19)|v3_pre_topc(X19,X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|(v2_struct_0(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t106_tmap_1])])])])])).
fof(c_0_32, plain, ![X22, X23, X24, X25]:(((~v5_pre_topc(X24,X22,X23)|(~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X23)))|(~v3_pre_topc(X25,X23)|v3_pre_topc(k8_relset_1(u1_struct_0(X22),u1_struct_0(X23),X24,X25),X22)))|k2_struct_0(X23)=k1_xboole_0|(~v1_funct_1(X24)|(~v1_funct_2(X24,u1_struct_0(X22),u1_struct_0(X23))|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(X23))))))|~l1_pre_topc(X23)|~l1_pre_topc(X22))&((m1_subset_1(esk3_3(X22,X23,X24),k1_zfmisc_1(u1_struct_0(X23)))|v5_pre_topc(X24,X22,X23)|k2_struct_0(X23)=k1_xboole_0|(~v1_funct_1(X24)|(~v1_funct_2(X24,u1_struct_0(X22),u1_struct_0(X23))|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(X23))))))|~l1_pre_topc(X23)|~l1_pre_topc(X22))&((v3_pre_topc(esk3_3(X22,X23,X24),X23)|v5_pre_topc(X24,X22,X23)|k2_struct_0(X23)=k1_xboole_0|(~v1_funct_1(X24)|(~v1_funct_2(X24,u1_struct_0(X22),u1_struct_0(X23))|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(X23))))))|~l1_pre_topc(X23)|~l1_pre_topc(X22))&(~v3_pre_topc(k8_relset_1(u1_struct_0(X22),u1_struct_0(X23),X24,esk3_3(X22,X23,X24)),X22)|v5_pre_topc(X24,X22,X23)|k2_struct_0(X23)=k1_xboole_0|(~v1_funct_1(X24)|(~v1_funct_2(X24,u1_struct_0(X22),u1_struct_0(X23))|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(X23))))))|~l1_pre_topc(X23)|~l1_pre_topc(X22)))))&((~v5_pre_topc(X24,X22,X23)|(~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X23)))|(~v3_pre_topc(X25,X23)|v3_pre_topc(k8_relset_1(u1_struct_0(X22),u1_struct_0(X23),X24,X25),X22)))|k2_struct_0(X22)!=k1_xboole_0|(~v1_funct_1(X24)|(~v1_funct_2(X24,u1_struct_0(X22),u1_struct_0(X23))|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(X23))))))|~l1_pre_topc(X23)|~l1_pre_topc(X22))&((m1_subset_1(esk3_3(X22,X23,X24),k1_zfmisc_1(u1_struct_0(X23)))|v5_pre_topc(X24,X22,X23)|k2_struct_0(X22)!=k1_xboole_0|(~v1_funct_1(X24)|(~v1_funct_2(X24,u1_struct_0(X22),u1_struct_0(X23))|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(X23))))))|~l1_pre_topc(X23)|~l1_pre_topc(X22))&((v3_pre_topc(esk3_3(X22,X23,X24),X23)|v5_pre_topc(X24,X22,X23)|k2_struct_0(X22)!=k1_xboole_0|(~v1_funct_1(X24)|(~v1_funct_2(X24,u1_struct_0(X22),u1_struct_0(X23))|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(X23))))))|~l1_pre_topc(X23)|~l1_pre_topc(X22))&(~v3_pre_topc(k8_relset_1(u1_struct_0(X22),u1_struct_0(X23),X24,esk3_3(X22,X23,X24)),X22)|v5_pre_topc(X24,X22,X23)|k2_struct_0(X22)!=k1_xboole_0|(~v1_funct_1(X24)|(~v1_funct_2(X24,u1_struct_0(X22),u1_struct_0(X23))|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(X23))))))|~l1_pre_topc(X23)|~l1_pre_topc(X22)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_tops_2])])])])])).
cnf(c_0_33, plain, (m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_34, plain, ![X14, X15]:(~m1_subset_1(X15,k1_zfmisc_1(X14))|k8_relset_1(X14,X14,k6_partfun1(X14),X15)=X15), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t171_funct_2])])).
cnf(c_0_35, plain, (v2_struct_0(X1)|k7_tmap_1(X1,X2)=k6_partfun1(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_36, plain, (l1_pre_topc(k6_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_37, plain, (v2_struct_0(X1)|v3_pre_topc(X3,k6_tmap_1(X1,X2))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k6_tmap_1(X1,X2))))|X3!=X2), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_38, negated_conjecture, (~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v3_pre_topc(esk2_0,esk1_0)|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_28]), c_0_28])).
cnf(c_0_39, negated_conjecture, (v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_16]), c_0_28]), c_0_17]), c_0_18])]), c_0_19])).
cnf(c_0_40, plain, (v5_pre_topc(X4,X2,g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3)))|~v5_pre_topc(X1,X2,X3)|X1!=X4|~v1_funct_1(X4)|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))))))|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_41, plain, (g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))=k6_tmap_1(X2,X1)|v2_struct_0(X2)|~v3_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_42, plain, (v5_pre_topc(X3,X1,X2)|k2_struct_0(X2)=k1_xboole_0|~v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk3_3(X1,X2,X3)),X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43, negated_conjecture, (m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_16]), c_0_28]), c_0_17]), c_0_18])]), c_0_19])).
cnf(c_0_44, plain, (k8_relset_1(X2,X2,k6_partfun1(X2),X1)=X1|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45, negated_conjecture, (k6_partfun1(u1_struct_0(esk1_0))=k7_tmap_1(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_16]), c_0_17]), c_0_18])]), c_0_19])).
cnf(c_0_46, plain, (v3_pre_topc(esk3_3(X1,X2,X3),X2)|v5_pre_topc(X3,X1,X2)|k2_struct_0(X2)=k1_xboole_0|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_47, plain, (v3_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(X3),X1,X4),X2)|k2_struct_0(X3)=k1_xboole_0|~v5_pre_topc(X1,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~v3_pre_topc(X4,X3)|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(X3)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_48, negated_conjecture, (l1_pre_topc(k6_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_16]), c_0_17]), c_0_18])]), c_0_19])).
cnf(c_0_49, plain, (v3_pre_topc(X1,k6_tmap_1(X2,X1))|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k6_tmap_1(X2,X1))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_50, negated_conjecture, (~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~v3_pre_topc(esk2_0,esk1_0)|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_39])])).
cnf(c_0_51, plain, (v5_pre_topc(X1,X2,g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3)))|~v5_pre_topc(X1,X2,X3)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~v2_pre_topc(X2)), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_52, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))=k6_tmap_1(esk1_0,esk2_0)|~v3_pre_topc(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_16]), c_0_17]), c_0_18])]), c_0_19])).
cnf(c_0_53, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,esk1_0)|~v3_pre_topc(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,esk2_0),esk3_3(esk1_0,esk1_0,k7_tmap_1(esk1_0,esk2_0))),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_39]), c_0_22]), c_0_43]), c_0_17])])).
cnf(c_0_54, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,esk2_0),X1)=X1|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_55, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,esk1_0)|v3_pre_topc(esk3_3(esk1_0,esk1_0,k7_tmap_1(esk1_0,esk2_0)),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_39]), c_0_22]), c_0_43]), c_0_17])])).
cnf(c_0_56, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(esk1_0),X2,X3),X1)|~v5_pre_topc(X2,X1,k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2)|~v3_pre_topc(X3,k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_28]), c_0_48])])).
cnf(c_0_57, negated_conjecture, (v3_pre_topc(esk2_0,k6_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_16]), c_0_28]), c_0_16]), c_0_17]), c_0_18])]), c_0_19])).
fof(c_0_58, plain, ![X32]:(~l1_struct_0(X32)|k2_struct_0(X32)=u1_struct_0(X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
fof(c_0_59, plain, ![X31]:(~l1_pre_topc(X31)|l1_struct_0(X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_60, negated_conjecture, (~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~v3_pre_topc(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_43])])).
cnf(c_0_61, negated_conjecture, (v5_pre_topc(X1,X2,k6_tmap_1(esk1_0,esk2_0))|~v5_pre_topc(X1,X2,esk1_0)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~v3_pre_topc(esk2_0,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_28]), c_0_28]), c_0_17]), c_0_18])])).
cnf(c_0_62, plain, (m1_subset_1(esk3_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|v5_pre_topc(X3,X1,X2)|k2_struct_0(X2)=k1_xboole_0|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_63, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,esk1_0)|~m1_subset_1(esk3_3(esk1_0,esk1_0,k7_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55])).
cnf(c_0_64, plain, (v3_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(X3),X1,X4),X2)|~v5_pre_topc(X1,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~v3_pre_topc(X4,X3)|k2_struct_0(X2)!=k1_xboole_0|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(X3)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_65, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(esk1_0),X2,esk2_0),X1)|~v5_pre_topc(X2,X1,k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_16]), c_0_57])])).
cnf(c_0_66, negated_conjecture, (v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|v3_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_67, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_68, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_69, negated_conjecture, (~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,esk1_0)|~v3_pre_topc(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_39]), c_0_22]), c_0_43]), c_0_17]), c_0_18])])).
cnf(c_0_70, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_39]), c_0_22]), c_0_43]), c_0_17])]), c_0_63])).
cnf(c_0_71, negated_conjecture, (v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(esk1_0),X2,X3),X1)|k2_struct_0(X1)!=k1_xboole_0|~v5_pre_topc(X2,X1,k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2)|~v3_pre_topc(X3,k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_28]), c_0_48])])).
cnf(c_0_72, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v3_pre_topc(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,esk2_0),esk2_0),esk1_0)|v3_pre_topc(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_39]), c_0_22]), c_0_43]), c_0_17])])).
cnf(c_0_73, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_74, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|~v3_pre_topc(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_75, negated_conjecture, (v3_pre_topc(X1,esk1_0)|k2_struct_0(esk1_0)!=k1_xboole_0|~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~v3_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_44]), c_0_45]), c_0_45]), c_0_39]), c_0_45]), c_0_22]), c_0_45]), c_0_43]), c_0_17])])).
cnf(c_0_76, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v3_pre_topc(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_54]), c_0_16])])).
cnf(c_0_77, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|~v3_pre_topc(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_17])])).
cnf(c_0_78, negated_conjecture, (v3_pre_topc(X1,esk1_0)|u1_struct_0(esk1_0)!=k1_xboole_0|~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~v3_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_73]), c_0_17])])).
cnf(c_0_79, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_76]), c_0_28]), c_0_48])]), c_0_77])).
cnf(c_0_80, negated_conjecture, (v3_pre_topc(X1,esk1_0)|~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~v3_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_79]), c_0_79])])).
cnf(c_0_81, plain, (v3_pre_topc(esk3_3(X1,X2,X3),X2)|v5_pre_topc(X3,X1,X2)|k2_struct_0(X1)!=k1_xboole_0|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_82, negated_conjecture, (v3_pre_topc(esk2_0,esk1_0)|v3_pre_topc(X1,esk1_0)|~v3_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_80, c_0_66])).
cnf(c_0_83, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_xboole_0))), inference(rw,[status(thm)],[c_0_16, c_0_79])).
cnf(c_0_84, plain, (m1_subset_1(esk3_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|v5_pre_topc(X3,X1,X2)|k2_struct_0(X1)!=k1_xboole_0|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_85, negated_conjecture, (v5_pre_topc(X1,esk1_0,X2)|v3_pre_topc(esk3_3(esk1_0,X2,X1),X2)|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(X2))|~v1_funct_1(X1)|~v3_pre_topc(esk2_0,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X2))))|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_74]), c_0_17])])).
cnf(c_0_86, negated_conjecture, (v3_pre_topc(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_57])])).
cnf(c_0_87, negated_conjecture, (v5_pre_topc(X1,esk1_0,X2)|m1_subset_1(esk3_3(esk1_0,X2,X1),k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(X2))|~v1_funct_1(X1)|~v3_pre_topc(esk2_0,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X2))))|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_74]), c_0_17])])).
cnf(c_0_88, plain, (v5_pre_topc(X3,X1,X2)|~v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk3_3(X1,X2,X3)),X1)|k2_struct_0(X1)!=k1_xboole_0|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_89, negated_conjecture, (v5_pre_topc(X1,esk1_0,X2)|v3_pre_topc(esk3_3(esk1_0,X2,X1),X2)|~v1_funct_2(X1,k1_xboole_0,u1_struct_0(X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,u1_struct_0(X2))))|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_79]), c_0_79]), c_0_86])])).
cnf(c_0_90, negated_conjecture, (v5_pre_topc(X1,esk1_0,X2)|m1_subset_1(esk3_3(esk1_0,X2,X1),k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X1,k1_xboole_0,u1_struct_0(X2))|~v1_funct_1(X1)|~v3_pre_topc(esk2_0,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,u1_struct_0(X2))))|~l1_pre_topc(X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_79]), c_0_79])).
cnf(c_0_91, negated_conjecture, (v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,esk1_0)|k2_struct_0(esk1_0)!=k1_xboole_0|~v3_pre_topc(esk3_3(esk1_0,esk1_0,k7_tmap_1(esk1_0,esk2_0)),esk1_0)|~m1_subset_1(esk3_3(esk1_0,esk1_0,k7_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_54]), c_0_39]), c_0_22]), c_0_43]), c_0_17])])).
cnf(c_0_92, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_86])])).
cnf(c_0_93, negated_conjecture, (~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_86])])).
cnf(c_0_94, negated_conjecture, (v5_pre_topc(X1,esk1_0,esk1_0)|v3_pre_topc(esk3_3(esk1_0,esk1_0,X1),esk1_0)|~v1_funct_2(X1,k1_xboole_0,k1_xboole_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_79]), c_0_17])])).
cnf(c_0_95, negated_conjecture, (v1_funct_2(k7_tmap_1(esk1_0,esk2_0),k1_xboole_0,k1_xboole_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_79]), c_0_79])).
cnf(c_0_96, negated_conjecture, (m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_79]), c_0_79])).
cnf(c_0_97, negated_conjecture, (v5_pre_topc(X1,esk1_0,X2)|m1_subset_1(esk3_3(esk1_0,X2,X1),k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X1,k1_xboole_0,u1_struct_0(X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,u1_struct_0(X2))))|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_86])])).
cnf(c_0_98, negated_conjecture, (~v3_pre_topc(esk3_3(esk1_0,esk1_0,k7_tmap_1(esk1_0,esk2_0)),esk1_0)|~m1_subset_1(esk3_3(esk1_0,esk1_0,k7_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(k1_xboole_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_79]), c_0_92])]), c_0_93])).
cnf(c_0_99, negated_conjecture, (v3_pre_topc(esk3_3(esk1_0,esk1_0,k7_tmap_1(esk1_0,esk2_0)),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_22]), c_0_96])]), c_0_93])).
cnf(c_0_100, negated_conjecture, (v5_pre_topc(X1,esk1_0,esk1_0)|m1_subset_1(esk3_3(esk1_0,esk1_0,X1),k1_zfmisc_1(k1_xboole_0))|~v1_funct_2(X1,k1_xboole_0,k1_xboole_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_79]), c_0_17])])).
cnf(c_0_101, negated_conjecture, (~m1_subset_1(esk3_3(esk1_0,esk1_0,k7_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_98, c_0_99])])).
cnf(c_0_102, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_95]), c_0_22]), c_0_96])]), c_0_93]), c_0_101]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 103
# Proof object clause steps            : 78
# Proof object formula steps           : 25
# Proof object conjectures             : 58
# Proof object clause conjectures      : 55
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 35
# Proof object simplifying inferences  : 151
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 34
# Processed clauses                    : 572
# ...of these trivial                  : 6
# ...subsumed                          : 245
# ...remaining for further processing  : 321
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 14
# Backward-rewritten                   : 109
# Generated clauses                    : 898
# ...of the previous two non-trivial   : 931
# Contextual simplify-reflections      : 44
# Paramodulations                      : 893
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 161
#    Positive orientable unit clauses  : 18
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 139
# Current number of unprocessed clauses: 361
# ...number of literals in the above   : 2907
# Current number of archived formulas  : 0
# Current number of archived clauses   : 157
# Clause-clause subsumption calls (NU) : 9197
# Rec. Clause-clause subsumption calls : 2774
# Non-unit clause-clause subsumptions  : 289
# Unit Clause-clause subsumption calls : 215
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 42861

# -------------------------------------------------
# User time                : 0.035 s
# System time              : 0.000 s
# Total time               : 0.035 s
# Maximum resident set size: 28752 pages
