# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t13_finset_1, axiom, ![X1, X2]:((r1_tarski(X1,X2)&v1_finset_1(X2))=>v1_finset_1(X1)), file('data/MPTP/proofs/11e208835a4d46959ec045474e189270.E_input', t13_finset_1)).
fof(t100_zfmisc_1, axiom, ![X1]:r1_tarski(X1,k1_zfmisc_1(k3_tarski(X1))), file('data/MPTP/proofs/11e208835a4d46959ec045474e189270.E_input', t100_zfmisc_1)).
fof(fc17_finset_1, axiom, ![X1]:(v1_finset_1(X1)=>v1_finset_1(k1_zfmisc_1(X1))), file('data/MPTP/proofs/11e208835a4d46959ec045474e189270.E_input', fc17_finset_1)).
fof(t25_finset_1, conjecture, ![X1]:((v1_finset_1(X1)&![X2]:(r2_hidden(X2,X1)=>v1_finset_1(X2)))<=>v1_finset_1(k3_tarski(X1))), file('data/MPTP/proofs/11e208835a4d46959ec045474e189270.E_input', t25_finset_1)).
fof(l22_finset_1, axiom, ![X1]:((v1_finset_1(X1)&![X2]:(r2_hidden(X2,X1)=>v1_finset_1(X2)))=>v1_finset_1(k3_tarski(X1))), file('data/MPTP/proofs/11e208835a4d46959ec045474e189270.E_input', l22_finset_1)).
fof(l49_zfmisc_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>r1_tarski(X1,k3_tarski(X2))), file('data/MPTP/proofs/11e208835a4d46959ec045474e189270.E_input', l49_zfmisc_1)).
fof(c_0_6, plain, ![X11, X12]:(~r1_tarski(X11,X12)|~v1_finset_1(X12)|v1_finset_1(X11)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_finset_1])])).
fof(c_0_7, plain, ![X13]:r1_tarski(X13,k1_zfmisc_1(k3_tarski(X13))), inference(variable_rename,[status(thm)],[t100_zfmisc_1])).
cnf(c_0_8, plain, (v1_finset_1(X1)|~r1_tarski(X1,X2)|~v1_finset_1(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, plain, (r1_tarski(X1,k1_zfmisc_1(k3_tarski(X1)))), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_10, plain, ![X10]:(~v1_finset_1(X10)|v1_finset_1(k1_zfmisc_1(X10))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc17_finset_1])])).
fof(c_0_11, negated_conjecture, ~(![X1]:((v1_finset_1(X1)&![X2]:(r2_hidden(X2,X1)=>v1_finset_1(X2)))<=>v1_finset_1(k3_tarski(X1)))), inference(assume_negation,[status(cth)],[t25_finset_1])).
cnf(c_0_12, plain, (v1_finset_1(X1)|~v1_finset_1(k1_zfmisc_1(k3_tarski(X1)))), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_13, plain, (v1_finset_1(k1_zfmisc_1(X1))|~v1_finset_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_14, negated_conjecture, ![X5]:(((r2_hidden(esk2_0,esk1_0)|~v1_finset_1(esk1_0)|~v1_finset_1(k3_tarski(esk1_0)))&(~v1_finset_1(esk2_0)|~v1_finset_1(esk1_0)|~v1_finset_1(k3_tarski(esk1_0))))&((v1_finset_1(esk1_0)|v1_finset_1(k3_tarski(esk1_0)))&(~r2_hidden(X5,esk1_0)|v1_finset_1(X5)|v1_finset_1(k3_tarski(esk1_0))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).
fof(c_0_15, plain, ![X8]:((r2_hidden(esk3_1(X8),X8)|~v1_finset_1(X8)|v1_finset_1(k3_tarski(X8)))&(~v1_finset_1(esk3_1(X8))|~v1_finset_1(X8)|v1_finset_1(k3_tarski(X8)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l22_finset_1])])])])).
cnf(c_0_16, plain, (v1_finset_1(X1)|~v1_finset_1(k3_tarski(X1))), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_17, negated_conjecture, (v1_finset_1(esk1_0)|v1_finset_1(k3_tarski(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, negated_conjecture, (v1_finset_1(X1)|v1_finset_1(k3_tarski(esk1_0))|~r2_hidden(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (r2_hidden(esk3_1(X1),X1)|v1_finset_1(k3_tarski(X1))|~v1_finset_1(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, negated_conjecture, (v1_finset_1(esk1_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
fof(c_0_21, plain, ![X6, X7]:(~r2_hidden(X6,X7)|r1_tarski(X6,k3_tarski(X7))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l49_zfmisc_1])])).
cnf(c_0_22, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_finset_1(esk1_0)|~v1_finset_1(k3_tarski(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23, plain, (v1_finset_1(k3_tarski(X1))|~v1_finset_1(esk3_1(X1))|~v1_finset_1(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, negated_conjecture, (v1_finset_1(esk3_1(esk1_0))|v1_finset_1(k3_tarski(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])])).
cnf(c_0_25, negated_conjecture, (~v1_finset_1(esk2_0)|~v1_finset_1(esk1_0)|~v1_finset_1(k3_tarski(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_26, plain, (r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_finset_1(k3_tarski(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_20])])).
cnf(c_0_28, negated_conjecture, (v1_finset_1(k3_tarski(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_20])])).
cnf(c_0_29, negated_conjecture, (~v1_finset_1(k3_tarski(esk1_0))|~v1_finset_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_20])])).
cnf(c_0_30, plain, (v1_finset_1(X1)|~r2_hidden(X1,X2)|~v1_finset_1(k3_tarski(X2))), inference(spm,[status(thm)],[c_0_8, c_0_26])).
cnf(c_0_31, negated_conjecture, (r2_hidden(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_28])])).
cnf(c_0_32, negated_conjecture, (~v1_finset_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_28])])).
cnf(c_0_33, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_28])]), c_0_32]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 34
# Proof object clause steps            : 21
# Proof object formula steps           : 13
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 15
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 10
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 10
# Processed clauses                    : 30
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 30
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 7
# Generated clauses                    : 12
# ...of the previous two non-trivial   : 10
# Contextual simplify-reflections      : 0
# Paramodulations                      : 12
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 13
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 8
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 17
# Clause-clause subsumption calls (NU) : 12
# Rec. Clause-clause subsumption calls : 9
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 841

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.000 s
# Total time               : 0.009 s
# Maximum resident set size: 28788 pages
