# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/11ecce4345b5493592c394b51f11dd7b.E_input', d3_tarski)).
fof(t44_tops_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>r1_tarski(k1_tops_1(X1,X2),X2))), file('data/MPTP/proofs/11ecce4345b5493592c394b51f11dd7b.E_input', t44_tops_1)).
fof(d1_connsp_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(m1_connsp_2(X3,X1,X2)<=>r2_hidden(X2,k1_tops_1(X1,X3)))))), file('data/MPTP/proofs/11ecce4345b5493592c394b51f11dd7b.E_input', d1_connsp_2)).
fof(t9_connsp_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~((r2_hidden(X3,X2)&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>~((m1_connsp_2(X4,X1,X3)&r1_tarski(X4,X2)))))))))), file('data/MPTP/proofs/11ecce4345b5493592c394b51f11dd7b.E_input', t9_connsp_2)).
fof(t7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,X2)=>r2_hidden(X4,X3)))=>r1_tarski(X2,X3)))), file('data/MPTP/proofs/11ecce4345b5493592c394b51f11dd7b.E_input', t7_subset_1)).
fof(dt_k1_tops_1, axiom, ![X1, X2]:((l1_pre_topc(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k1_tops_1(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/11ecce4345b5493592c394b51f11dd7b.E_input', dt_k1_tops_1)).
fof(t48_tops_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r1_tarski(X2,X3)=>r1_tarski(k1_tops_1(X1,X2),k1_tops_1(X1,X3)))))), file('data/MPTP/proofs/11ecce4345b5493592c394b51f11dd7b.E_input', t48_tops_1)).
fof(projectivity_k1_tops_1, axiom, ![X1, X2]:((l1_pre_topc(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>k1_tops_1(X1,k1_tops_1(X1,X2))=k1_tops_1(X1,X2)), file('data/MPTP/proofs/11ecce4345b5493592c394b51f11dd7b.E_input', projectivity_k1_tops_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/11ecce4345b5493592c394b51f11dd7b.E_input', d10_xboole_0)).
fof(t5_connsp_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((v3_pre_topc(X2,X1)&r2_hidden(X3,X2))=>m1_connsp_2(X2,X1,X3))))), file('data/MPTP/proofs/11ecce4345b5493592c394b51f11dd7b.E_input', t5_connsp_2)).
fof(t55_tops_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(l1_pre_topc(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>((v3_pre_topc(X4,X2)=>k1_tops_1(X2,X4)=X4)&(k1_tops_1(X1,X3)=X3=>v3_pre_topc(X3,X1))))))), file('data/MPTP/proofs/11ecce4345b5493592c394b51f11dd7b.E_input', t55_tops_1)).
fof(c_0_11, plain, ![X5, X6, X7, X8, X9]:((~r1_tarski(X5,X6)|(~r2_hidden(X7,X5)|r2_hidden(X7,X6)))&((r2_hidden(esk1_2(X8,X9),X8)|r1_tarski(X8,X9))&(~r2_hidden(esk1_2(X8,X9),X9)|r1_tarski(X8,X9)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_12, plain, ![X11, X12]:(~l1_pre_topc(X11)|(~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|r1_tarski(k1_tops_1(X11,X12),X12))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t44_tops_1])])])).
cnf(c_0_13, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_14, plain, (r1_tarski(k1_tops_1(X1,X2),X2)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_15, plain, ![X13, X14, X15]:((~m1_connsp_2(X15,X13,X14)|r2_hidden(X14,k1_tops_1(X13,X15))|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X13)))|~m1_subset_1(X14,u1_struct_0(X13))|(v2_struct_0(X13)|(~v2_pre_topc(X13)|~l1_pre_topc(X13))))&(~r2_hidden(X14,k1_tops_1(X13,X15))|m1_connsp_2(X15,X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X13)))|~m1_subset_1(X14,u1_struct_0(X13))|(v2_struct_0(X13)|(~v2_pre_topc(X13)|~l1_pre_topc(X13))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_connsp_2])])])])])).
fof(c_0_16, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~((r2_hidden(X3,X2)&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>~((m1_connsp_2(X4,X1,X3)&r1_tarski(X4,X2))))))))))), inference(assume_negation,[status(cth)],[t9_connsp_2])).
cnf(c_0_17, plain, (r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3)|~r2_hidden(X1,k1_tops_1(X3,X2))), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_18, plain, (r2_hidden(X3,k1_tops_1(X2,X1))|v2_struct_0(X2)|~m1_connsp_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,u1_struct_0(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_19, negated_conjecture, ![X19, X20]:((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&(((m1_subset_1(esk4_0,u1_struct_0(esk2_0))|~v3_pre_topc(esk3_0,esk2_0))&((r2_hidden(esk4_0,esk3_0)|~v3_pre_topc(esk3_0,esk2_0))&(~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(esk2_0)))|(~m1_connsp_2(X19,esk2_0,esk4_0)|~r1_tarski(X19,esk3_0))|~v3_pre_topc(esk3_0,esk2_0))))&((m1_subset_1(esk5_1(X20),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X20,esk3_0)|~m1_subset_1(X20,u1_struct_0(esk2_0))|v3_pre_topc(esk3_0,esk2_0))&((m1_connsp_2(esk5_1(X20),esk2_0,X20)|~r2_hidden(X20,esk3_0)|~m1_subset_1(X20,u1_struct_0(esk2_0))|v3_pre_topc(esk3_0,esk2_0))&(r1_tarski(esk5_1(X20),esk3_0)|~r2_hidden(X20,esk3_0)|~m1_subset_1(X20,u1_struct_0(esk2_0))|v3_pre_topc(esk3_0,esk2_0))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_16])])])])])])).
fof(c_0_20, plain, ![X22, X23, X24]:((m1_subset_1(esk6_3(X22,X23,X24),X22)|r1_tarski(X23,X24)|~m1_subset_1(X24,k1_zfmisc_1(X22))|~m1_subset_1(X23,k1_zfmisc_1(X22)))&((r2_hidden(esk6_3(X22,X23,X24),X23)|r1_tarski(X23,X24)|~m1_subset_1(X24,k1_zfmisc_1(X22))|~m1_subset_1(X23,k1_zfmisc_1(X22)))&(~r2_hidden(esk6_3(X22,X23,X24),X24)|r1_tarski(X23,X24)|~m1_subset_1(X24,k1_zfmisc_1(X22))|~m1_subset_1(X23,k1_zfmisc_1(X22))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_subset_1])])])])])).
cnf(c_0_21, plain, (v2_struct_0(X1)|r2_hidden(X2,X3)|~m1_connsp_2(X3,X1,X2)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_22, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23, negated_conjecture, (v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (m1_connsp_2(X3,X2,X1)|v2_struct_0(X2)|~r2_hidden(X1,k1_tops_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_27, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r1_tarski(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, negated_conjecture, (r2_hidden(X1,esk3_0)|~m1_connsp_2(esk3_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24])]), c_0_25])).
cnf(c_0_29, plain, (m1_connsp_2(X1,X2,esk6_3(X3,k1_tops_1(X2,X1),X4))|v2_struct_0(X2)|r1_tarski(k1_tops_1(X2,X1),X4)|~v2_pre_topc(X2)|~m1_subset_1(esk6_3(X3,k1_tops_1(X2,X1),X4),u1_struct_0(X2))|~m1_subset_1(k1_tops_1(X2,X1),k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(X3))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_30, negated_conjecture, (r2_hidden(esk6_3(X1,k1_tops_1(esk2_0,esk3_0),X2),esk3_0)|r1_tarski(k1_tops_1(esk2_0,esk3_0),X2)|~m1_subset_1(esk6_3(X1,k1_tops_1(esk2_0,esk3_0),X2),u1_struct_0(esk2_0))|~m1_subset_1(k1_tops_1(esk2_0,esk3_0),k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_23]), c_0_22]), c_0_24])]), c_0_25])).
cnf(c_0_31, plain, (m1_subset_1(esk6_3(X1,X2,X3),X1)|r1_tarski(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_32, plain, (r1_tarski(X2,X3)|~r2_hidden(esk6_3(X1,X2,X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_33, negated_conjecture, (r2_hidden(esk6_3(u1_struct_0(esk2_0),k1_tops_1(esk2_0,esk3_0),X1),esk3_0)|r1_tarski(k1_tops_1(esk2_0,esk3_0),X1)|~m1_subset_1(k1_tops_1(esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
fof(c_0_34, plain, ![X26, X27]:(~l1_pre_topc(X26)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X26)))|m1_subset_1(k1_tops_1(X26,X27),k1_zfmisc_1(u1_struct_0(X26)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_tops_1])])).
fof(c_0_35, plain, ![X28, X29, X30]:(~l1_pre_topc(X28)|(~m1_subset_1(X29,k1_zfmisc_1(u1_struct_0(X28)))|(~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X28)))|(~r1_tarski(X29,X30)|r1_tarski(k1_tops_1(X28,X29),k1_tops_1(X28,X30)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t48_tops_1])])])).
cnf(c_0_36, negated_conjecture, (r1_tarski(k1_tops_1(esk2_0,esk3_0),esk3_0)|~m1_subset_1(k1_tops_1(esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_22])])).
cnf(c_0_37, plain, (m1_subset_1(k1_tops_1(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_38, plain, (r1_tarski(k1_tops_1(X1,X2),k1_tops_1(X1,X3))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_39, negated_conjecture, (r1_tarski(k1_tops_1(esk2_0,esk3_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_22]), c_0_24])])).
cnf(c_0_40, plain, (r2_hidden(X1,k1_tops_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~r2_hidden(X1,k1_tops_1(X2,X4))|~r1_tarski(X4,X3)), inference(spm,[status(thm)],[c_0_13, c_0_38])).
cnf(c_0_41, plain, (r2_hidden(esk1_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_42, plain, ![X31, X32]:(~l1_pre_topc(X31)|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|k1_tops_1(X31,k1_tops_1(X31,X32))=k1_tops_1(X31,X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[projectivity_k1_tops_1])])).
cnf(c_0_43, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,k1_tops_1(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_13, c_0_39])).
cnf(c_0_44, plain, (r2_hidden(esk1_2(k1_tops_1(X1,X2),X3),k1_tops_1(X1,X4))|r1_tarski(k1_tops_1(X1,X2),X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~r1_tarski(X2,X4)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_45, plain, (k1_tops_1(X1,k1_tops_1(X1,X2))=k1_tops_1(X1,X2)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_46, plain, (r1_tarski(X1,X2)|~r2_hidden(esk1_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_47, negated_conjecture, (r2_hidden(esk1_2(k1_tops_1(esk2_0,X1),X2),esk3_0)|r1_tarski(k1_tops_1(esk2_0,X1),X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r1_tarski(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_22]), c_0_24])])).
fof(c_0_48, plain, ![X33, X34]:(((r1_tarski(X33,X34)|X33!=X34)&(r1_tarski(X34,X33)|X33!=X34))&(~r1_tarski(X33,X34)|~r1_tarski(X34,X33)|X33=X34)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_49, plain, (r2_hidden(X1,k1_tops_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~r2_hidden(X1,k1_tops_1(X2,X4))|~r1_tarski(k1_tops_1(X2,X4),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_45]), c_0_37])).
cnf(c_0_50, negated_conjecture, (r1_tarski(k1_tops_1(esk2_0,X1),esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_51, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_52, plain, ![X35, X36, X37]:(v2_struct_0(X35)|(~v2_pre_topc(X35)|~l1_pre_topc(X35))|(~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X35)))|(~m1_subset_1(X37,u1_struct_0(X35))|(~v3_pre_topc(X36,X35)|~r2_hidden(X37,X36)|m1_connsp_2(X36,X35,X37))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t5_connsp_2])])])])).
cnf(c_0_53, negated_conjecture, (r2_hidden(X1,k1_tops_1(esk2_0,esk3_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,k1_tops_1(esk2_0,X2))|~r1_tarski(X2,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_22]), c_0_24])])).
cnf(c_0_54, negated_conjecture, (~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_connsp_2(X1,esk2_0,esk4_0)|~r1_tarski(X1,esk3_0)|~v3_pre_topc(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_55, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_51])).
cnf(c_0_56, plain, (v2_struct_0(X1)|m1_connsp_2(X2,X1,X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_pre_topc(X2,X1)|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_57, negated_conjecture, (r2_hidden(esk1_2(k1_tops_1(esk2_0,X1),X2),k1_tops_1(esk2_0,esk3_0))|r1_tarski(k1_tops_1(esk2_0,X1),X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_53, c_0_41])).
cnf(c_0_58, negated_conjecture, (~v3_pre_topc(esk3_0,esk2_0)|~m1_connsp_2(esk3_0,esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_22]), c_0_55])])).
cnf(c_0_59, negated_conjecture, (m1_connsp_2(esk3_0,esk2_0,X1)|~v3_pre_topc(esk3_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_22]), c_0_23]), c_0_24])]), c_0_25])).
cnf(c_0_60, negated_conjecture, (r2_hidden(esk4_0,esk3_0)|~v3_pre_topc(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_61, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk2_0))|~v3_pre_topc(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_62, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_63, negated_conjecture, (r1_tarski(k1_tops_1(esk2_0,X1),k1_tops_1(esk2_0,esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_46, c_0_57])).
cnf(c_0_64, negated_conjecture, (r1_tarski(esk5_1(X1),esk3_0)|v3_pre_topc(esk3_0,esk2_0)|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_65, negated_conjecture, (~v3_pre_topc(esk3_0,esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_60]), c_0_61])).
cnf(c_0_66, negated_conjecture, (k1_tops_1(esk2_0,esk3_0)=k1_tops_1(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r1_tarski(k1_tops_1(esk2_0,esk3_0),k1_tops_1(esk2_0,X1))|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_67, plain, (r1_tarski(k1_tops_1(X1,X2),k1_tops_1(X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~r1_tarski(k1_tops_1(X1,X2),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_45]), c_0_37])).
cnf(c_0_68, negated_conjecture, (r2_hidden(X1,esk3_0)|~m1_subset_1(X2,u1_struct_0(esk2_0))|~r2_hidden(X1,esk5_1(X2))|~r2_hidden(X2,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_64]), c_0_65])).
cnf(c_0_69, negated_conjecture, (k1_tops_1(esk2_0,esk3_0)=k1_tops_1(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r1_tarski(k1_tops_1(esk2_0,esk3_0),X1)|~r1_tarski(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_22]), c_0_24])])).
cnf(c_0_70, negated_conjecture, (r2_hidden(esk1_2(esk5_1(X1),X2),esk3_0)|r1_tarski(esk5_1(X1),X2)|~m1_subset_1(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_68, c_0_41])).
cnf(c_0_71, negated_conjecture, (k1_tops_1(esk2_0,k1_tops_1(esk2_0,X1))=k1_tops_1(esk2_0,esk3_0)|~m1_subset_1(k1_tops_1(esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r1_tarski(k1_tops_1(esk2_0,X1),esk3_0)|~r1_tarski(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_38]), c_0_22]), c_0_24])])).
cnf(c_0_72, negated_conjecture, (r2_hidden(X1,k1_tops_1(esk2_0,esk3_0))|~m1_connsp_2(X2,esk2_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~r1_tarski(X2,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_18]), c_0_23]), c_0_24])]), c_0_25])).
cnf(c_0_73, negated_conjecture, (m1_subset_1(esk5_1(X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|v3_pre_topc(esk3_0,esk2_0)|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_74, negated_conjecture, (r1_tarski(esk5_1(X1),esk3_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_46, c_0_70])).
cnf(c_0_75, negated_conjecture, (k1_tops_1(esk2_0,k1_tops_1(esk2_0,X1))=k1_tops_1(esk2_0,esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r1_tarski(k1_tops_1(esk2_0,X1),esk3_0)|~r1_tarski(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_37]), c_0_24])])).
fof(c_0_76, plain, ![X38, X39, X40, X41]:((~v3_pre_topc(X41,X39)|k1_tops_1(X39,X41)=X41|~m1_subset_1(X41,k1_zfmisc_1(u1_struct_0(X39)))|~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X38)))|~l1_pre_topc(X39)|(~v2_pre_topc(X38)|~l1_pre_topc(X38)))&(k1_tops_1(X38,X40)!=X40|v3_pre_topc(X40,X38)|~m1_subset_1(X41,k1_zfmisc_1(u1_struct_0(X39)))|~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X38)))|~l1_pre_topc(X39)|(~v2_pre_topc(X38)|~l1_pre_topc(X38)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_tops_1])])])])).
cnf(c_0_77, negated_conjecture, (r2_hidden(X1,k1_tops_1(esk2_0,esk3_0))|~m1_connsp_2(esk5_1(X2),esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~r2_hidden(X2,esk3_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_65]), c_0_74])).
cnf(c_0_78, negated_conjecture, (m1_connsp_2(esk5_1(X1),esk2_0,X1)|v3_pre_topc(esk3_0,esk2_0)|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_79, negated_conjecture, (m1_subset_1(k1_tops_1(esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(k1_tops_1(esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r1_tarski(k1_tops_1(esk2_0,X1),esk3_0)|~r1_tarski(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_75]), c_0_24])])).
cnf(c_0_80, plain, (v3_pre_topc(X2,X1)|k1_tops_1(X1,X2)!=X2|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X4)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X4)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_81, negated_conjecture, (r2_hidden(X1,k1_tops_1(esk2_0,esk3_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_65])).
cnf(c_0_82, negated_conjecture, (m1_subset_1(k1_tops_1(esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r1_tarski(k1_tops_1(esk2_0,X1),esk3_0)|~r1_tarski(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_37]), c_0_24])])).
cnf(c_0_83, negated_conjecture, (v3_pre_topc(X1,X2)|k1_tops_1(X2,X1)!=X1|~v2_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_22]), c_0_24])])).
cnf(c_0_84, negated_conjecture, (r1_tarski(X1,k1_tops_1(esk2_0,esk3_0))|~m1_subset_1(esk6_3(X2,X1,k1_tops_1(esk2_0,esk3_0)),u1_struct_0(esk2_0))|~m1_subset_1(k1_tops_1(esk2_0,esk3_0),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(esk6_3(X2,X1,k1_tops_1(esk2_0,esk3_0)),esk3_0)), inference(spm,[status(thm)],[c_0_32, c_0_81])).
cnf(c_0_85, negated_conjecture, (m1_subset_1(k1_tops_1(esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_14]), c_0_22]), c_0_55]), c_0_24])])).
cnf(c_0_86, negated_conjecture, (k1_tops_1(esk2_0,esk3_0)!=esk3_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_22]), c_0_23]), c_0_24])]), c_0_65])).
cnf(c_0_87, negated_conjecture, (r1_tarski(X1,k1_tops_1(esk2_0,esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(esk6_3(u1_struct_0(esk2_0),X1,k1_tops_1(esk2_0,esk3_0)),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_31]), c_0_85])])).
cnf(c_0_88, negated_conjecture, (~r1_tarski(esk3_0,k1_tops_1(esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_39]), c_0_86])).
cnf(c_0_89, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_27]), c_0_22]), c_0_85])]), c_0_88]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 90
# Proof object clause steps            : 67
# Proof object formula steps           : 23
# Proof object conjectures             : 46
# Proof object clause conjectures      : 43
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 40
# Proof object simplifying inferences  : 68
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 28
# Processed clauses                    : 2660
# ...of these trivial                  : 18
# ...subsumed                          : 1991
# ...remaining for further processing  : 651
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 139
# Backward-rewritten                   : 35
# Generated clauses                    : 8917
# ...of the previous two non-trivial   : 7761
# Contextual simplify-reflections      : 40
# Paramodulations                      : 8915
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 475
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 463
# Current number of unprocessed clauses: 4879
# ...number of literals in the above   : 42029
# Current number of archived formulas  : 0
# Current number of archived clauses   : 174
# Clause-clause subsumption calls (NU) : 115021
# Rec. Clause-clause subsumption calls : 7829
# Non-unit clause-clause subsumptions  : 1758
# Unit Clause-clause subsumption calls : 59
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 14
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 342062

# -------------------------------------------------
# User time                : 0.226 s
# System time              : 0.011 s
# Total time               : 0.237 s
# Maximum resident set size: 28532 pages
