# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/11ee756c5f8d47798e79f2b83efcd7a6.E_input', d5_relat_1)).
fof(t37_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>(k2_relat_1(X1)=k1_relat_1(k4_relat_1(X1))&k1_relat_1(X1)=k2_relat_1(k4_relat_1(X1)))), file('data/MPTP/proofs/11ee756c5f8d47798e79f2b83efcd7a6.E_input', t37_relat_1)).
fof(dt_k4_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>v1_relat_1(k4_relat_1(X1))), file('data/MPTP/proofs/11ee756c5f8d47798e79f2b83efcd7a6.E_input', dt_k4_relat_1)).
fof(d13_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(X3=k9_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(k4_tarski(X5,X4),X1)&r2_hidden(X5,X2))))), file('data/MPTP/proofs/11ee756c5f8d47798e79f2b83efcd7a6.E_input', d13_relat_1)).
fof(t146_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>k9_relat_1(X1,k1_relat_1(X1))=k2_relat_1(X1)), file('data/MPTP/proofs/11ee756c5f8d47798e79f2b83efcd7a6.E_input', t146_relat_1)).
fof(d7_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(X2=k4_relat_1(X1)<=>![X3, X4]:(r2_hidden(k4_tarski(X3,X4),X2)<=>r2_hidden(k4_tarski(X4,X3),X1))))), file('data/MPTP/proofs/11ee756c5f8d47798e79f2b83efcd7a6.E_input', d7_relat_1)).
fof(c_0_6, plain, ![X29, X30, X31, X33, X34, X35, X36, X38]:(((~r2_hidden(X31,X30)|r2_hidden(k4_tarski(esk7_3(X29,X30,X31),X31),X29)|X30!=k2_relat_1(X29))&(~r2_hidden(k4_tarski(X34,X33),X29)|r2_hidden(X33,X30)|X30!=k2_relat_1(X29)))&((~r2_hidden(esk8_2(X35,X36),X36)|~r2_hidden(k4_tarski(X38,esk8_2(X35,X36)),X35)|X36=k2_relat_1(X35))&(r2_hidden(esk8_2(X35,X36),X36)|r2_hidden(k4_tarski(esk9_2(X35,X36),esk8_2(X35,X36)),X35)|X36=k2_relat_1(X35)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
fof(c_0_7, plain, ![X7]:((k2_relat_1(X7)=k1_relat_1(k4_relat_1(X7))|~v1_relat_1(X7))&(k1_relat_1(X7)=k2_relat_1(k4_relat_1(X7))|~v1_relat_1(X7))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t37_relat_1])])])).
cnf(c_0_8, plain, (r2_hidden(k4_tarski(esk7_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, plain, (k1_relat_1(X1)=k2_relat_1(k4_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_10, plain, ![X28]:(~v1_relat_1(X28)|v1_relat_1(k4_relat_1(X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_relat_1])])).
cnf(c_0_11, plain, (r2_hidden(k4_tarski(esk7_3(k4_relat_1(X1),X2,X3),X3),k4_relat_1(X1))|X2!=k1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_12, plain, (k2_relat_1(X1)=k1_relat_1(k4_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, plain, (v1_relat_1(k4_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (r2_hidden(X2,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15, plain, (r2_hidden(k4_tarski(esk7_3(k4_relat_1(k4_relat_1(X1)),X2,X3),X3),k4_relat_1(k4_relat_1(X1)))|X2!=k2_relat_1(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])).
fof(c_0_16, plain, ![X8, X9, X10, X11, X13, X14, X15, X16, X18]:((((r2_hidden(k4_tarski(esk2_4(X8,X9,X10,X11),X11),X8)|~r2_hidden(X11,X10)|X10!=k9_relat_1(X8,X9)|~v1_relat_1(X8))&(r2_hidden(esk2_4(X8,X9,X10,X11),X9)|~r2_hidden(X11,X10)|X10!=k9_relat_1(X8,X9)|~v1_relat_1(X8)))&(~r2_hidden(k4_tarski(X14,X13),X8)|~r2_hidden(X14,X9)|r2_hidden(X13,X10)|X10!=k9_relat_1(X8,X9)|~v1_relat_1(X8)))&((~r2_hidden(esk3_3(X8,X15,X16),X16)|(~r2_hidden(k4_tarski(X18,esk3_3(X8,X15,X16)),X8)|~r2_hidden(X18,X15))|X16=k9_relat_1(X8,X15)|~v1_relat_1(X8))&((r2_hidden(k4_tarski(esk4_3(X8,X15,X16),esk3_3(X8,X15,X16)),X8)|r2_hidden(esk3_3(X8,X15,X16),X16)|X16=k9_relat_1(X8,X15)|~v1_relat_1(X8))&(r2_hidden(esk4_3(X8,X15,X16),X15)|r2_hidden(esk3_3(X8,X15,X16),X16)|X16=k9_relat_1(X8,X15)|~v1_relat_1(X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d13_relat_1])])])])])])).
cnf(c_0_17, plain, (r2_hidden(X1,X2)|X2!=k2_relat_1(k4_relat_1(k4_relat_1(X3)))|X4!=k2_relat_1(X3)|~r2_hidden(X1,X4)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_18, plain, (r2_hidden(esk8_2(X1,X2),X2)|r2_hidden(k4_tarski(esk9_2(X1,X2),esk8_2(X1,X2)),X1)|X2=k2_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_19, plain, (r2_hidden(k4_tarski(esk2_4(X1,X2,X3,X4),X4),X1)|~r2_hidden(X4,X3)|X3!=k9_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, (r2_hidden(X1,X2)|X2!=k1_relat_1(k4_relat_1(X3))|X4!=k2_relat_1(X3)|~r2_hidden(X1,X4)|~v1_relat_1(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_9]), c_0_13])).
cnf(c_0_21, plain, (X1=k2_relat_1(X2)|r2_hidden(esk8_2(X2,X1),X1)|r2_hidden(esk8_2(X2,X1),X3)|X3!=k2_relat_1(X2)), inference(spm,[status(thm)],[c_0_14, c_0_18])).
fof(c_0_22, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>k9_relat_1(X1,k1_relat_1(X1))=k2_relat_1(X1))), inference(assume_negation,[status(cth)],[t146_relat_1])).
cnf(c_0_23, plain, (X2=k2_relat_1(X1)|~r2_hidden(esk8_2(X1,X2),X2)|~r2_hidden(k4_tarski(X3,esk8_2(X1,X2)),X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_24, plain, (r2_hidden(k4_tarski(esk2_4(X1,X2,k9_relat_1(X1,X2),X3),X3),X1)|~r2_hidden(X3,k9_relat_1(X1,X2))|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (r2_hidden(X1,k1_relat_1(k4_relat_1(X2)))|X3!=k2_relat_1(X2)|~r2_hidden(X1,X3)|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (X1=k2_relat_1(X2)|r2_hidden(esk8_2(X2,X1),k2_relat_1(X2))|r2_hidden(esk8_2(X2,X1),X1)), inference(er,[status(thm)],[c_0_21])).
fof(c_0_27, negated_conjecture, (v1_relat_1(esk1_0)&k9_relat_1(esk1_0,k1_relat_1(esk1_0))!=k2_relat_1(esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])).
fof(c_0_28, plain, ![X20, X21, X22, X23, X24, X25]:(((~r2_hidden(k4_tarski(X22,X23),X21)|r2_hidden(k4_tarski(X23,X22),X20)|X21!=k4_relat_1(X20)|~v1_relat_1(X21)|~v1_relat_1(X20))&(~r2_hidden(k4_tarski(X25,X24),X20)|r2_hidden(k4_tarski(X24,X25),X21)|X21!=k4_relat_1(X20)|~v1_relat_1(X21)|~v1_relat_1(X20)))&((~r2_hidden(k4_tarski(esk5_2(X20,X21),esk6_2(X20,X21)),X21)|~r2_hidden(k4_tarski(esk6_2(X20,X21),esk5_2(X20,X21)),X20)|X21=k4_relat_1(X20)|~v1_relat_1(X21)|~v1_relat_1(X20))&(r2_hidden(k4_tarski(esk5_2(X20,X21),esk6_2(X20,X21)),X21)|r2_hidden(k4_tarski(esk6_2(X20,X21),esk5_2(X20,X21)),X20)|X21=k4_relat_1(X20)|~v1_relat_1(X21)|~v1_relat_1(X20)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_relat_1])])])])])])).
cnf(c_0_29, plain, (r2_hidden(k4_tarski(esk7_3(X1,k2_relat_1(X1),X2),X2),X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_8])).
cnf(c_0_30, plain, (X1=k2_relat_1(X2)|~r2_hidden(esk8_2(X2,X1),k9_relat_1(X2,X3))|~r2_hidden(esk8_2(X2,X1),X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_31, plain, (X1=k2_relat_1(X2)|r2_hidden(esk8_2(X2,X1),k1_relat_1(k4_relat_1(X3)))|r2_hidden(esk8_2(X2,X1),X1)|k2_relat_1(X2)!=k2_relat_1(X3)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_32, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, plain, (r2_hidden(k4_tarski(X2,X1),X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k4_relat_1(X3)|~v1_relat_1(X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34, plain, (X1=k2_relat_1(X2)|~r2_hidden(esk8_2(X2,X1),k2_relat_1(X2))|~r2_hidden(esk8_2(X2,X1),X1)), inference(spm,[status(thm)],[c_0_23, c_0_29])).
cnf(c_0_35, plain, (k9_relat_1(X1,X2)=k2_relat_1(X1)|r2_hidden(esk8_2(X1,k9_relat_1(X1,X2)),k2_relat_1(X1))|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_26]), c_0_26])).
cnf(c_0_36, negated_conjecture, (X1=k2_relat_1(X2)|r2_hidden(esk8_2(X2,X1),k1_relat_1(k4_relat_1(esk1_0)))|r2_hidden(esk8_2(X2,X1),X1)|k2_relat_1(X2)!=k2_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_37, plain, (r2_hidden(k4_tarski(X1,esk7_3(X2,k2_relat_1(X2),X1)),X3)|X3!=k4_relat_1(X2)|~r2_hidden(X1,k2_relat_1(X2))|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_33, c_0_29])).
cnf(c_0_38, plain, (k9_relat_1(X1,X2)=k2_relat_1(X1)|~r2_hidden(esk8_2(X1,k9_relat_1(X1,X2)),k9_relat_1(X1,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_39, negated_conjecture, (X1=k2_relat_1(esk1_0)|r2_hidden(esk8_2(esk1_0,X1),k1_relat_1(k4_relat_1(esk1_0)))|r2_hidden(esk8_2(esk1_0,X1),X1)), inference(er,[status(thm)],[c_0_36])).
cnf(c_0_40, plain, (r2_hidden(X2,X5)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,X4)|X5!=k9_relat_1(X3,X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_41, plain, (r2_hidden(esk7_3(X1,k2_relat_1(X1),X2),X3)|X3!=k2_relat_1(X4)|X4!=k4_relat_1(X1)|~r2_hidden(X2,k2_relat_1(X1))|~v1_relat_1(X4)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_14, c_0_37])).
cnf(c_0_42, negated_conjecture, (k9_relat_1(esk1_0,X1)=k2_relat_1(esk1_0)|r2_hidden(esk8_2(esk1_0,k9_relat_1(esk1_0,X1)),k1_relat_1(k4_relat_1(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_32])])).
cnf(c_0_43, plain, (r2_hidden(X1,X2)|X2!=k9_relat_1(X3,X4)|~r2_hidden(esk7_3(X3,k2_relat_1(X3),X1),X4)|~r2_hidden(X1,k2_relat_1(X3))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_40, c_0_29])).
cnf(c_0_44, plain, (r2_hidden(esk7_3(X1,k2_relat_1(X1),X2),k2_relat_1(X3))|X3!=k4_relat_1(X1)|~r2_hidden(X2,k2_relat_1(X1))|~v1_relat_1(X3)|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_45, negated_conjecture, (k9_relat_1(esk1_0,X1)=k2_relat_1(esk1_0)|r2_hidden(esk8_2(esk1_0,k9_relat_1(esk1_0,X1)),k2_relat_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_12]), c_0_32])])).
cnf(c_0_46, plain, (r2_hidden(X1,X2)|X2!=k9_relat_1(X3,k2_relat_1(X4))|X4!=k4_relat_1(X3)|~r2_hidden(X1,k2_relat_1(X3))|~v1_relat_1(X3)|~v1_relat_1(X4)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_47, negated_conjecture, (k9_relat_1(esk1_0,X1)=k2_relat_1(esk1_0)|~r2_hidden(esk8_2(esk1_0,k9_relat_1(esk1_0,X1)),k9_relat_1(esk1_0,X1))), inference(spm,[status(thm)],[c_0_34, c_0_45])).
cnf(c_0_48, plain, (r2_hidden(X1,k9_relat_1(X2,k2_relat_1(X3)))|X3!=k4_relat_1(X2)|~r2_hidden(X1,k2_relat_1(X2))|~v1_relat_1(X2)|~v1_relat_1(X3)), inference(er,[status(thm)],[c_0_46])).
cnf(c_0_49, negated_conjecture, (k9_relat_1(esk1_0,k2_relat_1(X1))=k2_relat_1(esk1_0)|X1!=k4_relat_1(esk1_0)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_32])]), c_0_45])).
cnf(c_0_50, negated_conjecture, (k9_relat_1(esk1_0,k1_relat_1(esk1_0))!=k2_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_51, negated_conjecture, (k9_relat_1(esk1_0,k1_relat_1(X1))=k2_relat_1(esk1_0)|k4_relat_1(X1)!=k4_relat_1(esk1_0)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_9]), c_0_13])).
cnf(c_0_52, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_32])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 53
# Proof object clause steps            : 40
# Proof object formula steps           : 13
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 28
# Proof object simplifying inferences  : 13
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 19
# Processed clauses                    : 556
# ...of these trivial                  : 4
# ...subsumed                          : 143
# ...remaining for further processing  : 409
# Other redundant clauses eliminated   : 497
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 12
# Backward-rewritten                   : 0
# Generated clauses                    : 10114
# ...of the previous two non-trivial   : 9505
# Contextual simplify-reflections      : 29
# Paramodulations                      : 9491
# Factorizations                       : 14
# Equation resolutions                 : 609
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 378
#    Positive orientable unit clauses  : 1
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 376
# Current number of unprocessed clauses: 8950
# ...number of literals in the above   : 66748
# Current number of archived formulas  : 0
# Current number of archived clauses   : 31
# Clause-clause subsumption calls (NU) : 44337
# Rec. Clause-clause subsumption calls : 2786
# Non-unit clause-clause subsumptions  : 184
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 327443

# -------------------------------------------------
# User time                : 0.191 s
# System time              : 0.010 s
# Total time               : 0.200 s
# Maximum resident set size: 28692 pages
