# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/1233d9f8775140faa45fde491e3963e7.E_input', d1_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/1233d9f8775140faa45fde491e3963e7.E_input', d3_tarski)).
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/1233d9f8775140faa45fde491e3963e7.E_input', d2_zfmisc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/1233d9f8775140faa45fde491e3963e7.E_input', d10_xboole_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/1233d9f8775140faa45fde491e3963e7.E_input', fc1_xboole_0)).
fof(fc3_zfmisc_1, axiom, ![X1, X2]:(v1_xboole_0(X2)=>v1_xboole_0(k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/1233d9f8775140faa45fde491e3963e7.E_input', fc3_zfmisc_1)).
fof(t36_mcart_1, conjecture, ![X1, X2, X3, X4, X5, X6]:(k3_zfmisc_1(X1,X2,X3)=k3_zfmisc_1(X4,X5,X6)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|(X1=X4&(X2=X5&X3=X6)))))), file('data/MPTP/proofs/1233d9f8775140faa45fde491e3963e7.E_input', t36_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/1233d9f8775140faa45fde491e3963e7.E_input', d3_zfmisc_1)).
fof(t117_zfmisc_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&((r1_tarski(k2_zfmisc_1(X2,X1),k2_zfmisc_1(X3,X1))|r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))&~(r1_tarski(X2,X3))))), file('data/MPTP/proofs/1233d9f8775140faa45fde491e3963e7.E_input', t117_zfmisc_1)).
fof(t135_zfmisc_1, axiom, ![X1, X2]:((r1_tarski(X1,k2_zfmisc_1(X1,X2))|r1_tarski(X1,k2_zfmisc_1(X2,X1)))=>X1=k1_xboole_0), file('data/MPTP/proofs/1233d9f8775140faa45fde491e3963e7.E_input', t135_zfmisc_1)).
fof(fc10_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))=>~(v1_xboole_0(k2_zfmisc_1(X1,X2)))), file('data/MPTP/proofs/1233d9f8775140faa45fde491e3963e7.E_input', fc10_subset_1)).
fof(t195_relat_1, axiom, ![X1, X2]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&~((k1_relat_1(k2_zfmisc_1(X1,X2))=X1&k2_relat_1(k2_zfmisc_1(X1,X2))=X2))))), file('data/MPTP/proofs/1233d9f8775140faa45fde491e3963e7.E_input', t195_relat_1)).
fof(c_0_12, plain, ![X7, X8, X9]:((~v1_xboole_0(X7)|~r2_hidden(X8,X7))&(r2_hidden(esk1_1(X9),X9)|v1_xboole_0(X9))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_13, plain, ![X11, X12, X13, X14, X15]:((~r1_tarski(X11,X12)|(~r2_hidden(X13,X11)|r2_hidden(X13,X12)))&((r2_hidden(esk2_2(X14,X15),X14)|r1_tarski(X14,X15))&(~r2_hidden(esk2_2(X14,X15),X15)|r1_tarski(X14,X15)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_14, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_16, plain, ![X17, X18, X19, X20, X23, X24, X25, X26, X27, X28, X30, X31]:((((r2_hidden(esk3_4(X17,X18,X19,X20),X17)|~r2_hidden(X20,X19)|X19!=k2_zfmisc_1(X17,X18))&((r2_hidden(esk4_4(X17,X18,X19,X20),X18)|~r2_hidden(X20,X19)|X19!=k2_zfmisc_1(X17,X18))&(X20=k4_tarski(esk3_4(X17,X18,X19,X20),esk4_4(X17,X18,X19,X20))|~r2_hidden(X20,X19)|X19!=k2_zfmisc_1(X17,X18))))&(~r2_hidden(X24,X17)|(~r2_hidden(X25,X18)|X23!=k4_tarski(X24,X25))|r2_hidden(X23,X19)|X19!=k2_zfmisc_1(X17,X18)))&((~r2_hidden(esk5_3(X26,X27,X28),X28)|(~r2_hidden(X30,X26)|(~r2_hidden(X31,X27)|esk5_3(X26,X27,X28)!=k4_tarski(X30,X31)))|X28=k2_zfmisc_1(X26,X27))&((r2_hidden(esk6_3(X26,X27,X28),X26)|r2_hidden(esk5_3(X26,X27,X28),X28)|X28=k2_zfmisc_1(X26,X27))&((r2_hidden(esk7_3(X26,X27,X28),X27)|r2_hidden(esk5_3(X26,X27,X28),X28)|X28=k2_zfmisc_1(X26,X27))&(esk5_3(X26,X27,X28)=k4_tarski(esk6_3(X26,X27,X28),esk7_3(X26,X27,X28))|r2_hidden(esk5_3(X26,X27,X28),X28)|X28=k2_zfmisc_1(X26,X27)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
fof(c_0_17, plain, ![X54, X55]:(((r1_tarski(X54,X55)|X54!=X55)&(r1_tarski(X55,X54)|X54!=X55))&(~r1_tarski(X54,X55)|~r1_tarski(X55,X54)|X54=X55)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_18, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_19, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
fof(c_0_20, plain, ![X34, X35]:(~v1_xboole_0(X35)|v1_xboole_0(k2_zfmisc_1(X34,X35))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc3_zfmisc_1])])).
fof(c_0_21, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6]:(k3_zfmisc_1(X1,X2,X3)=k3_zfmisc_1(X4,X5,X6)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|(X1=X4&(X2=X5&X3=X6))))))), inference(assume_negation,[status(cth)],[t36_mcart_1])).
cnf(c_0_22, plain, (r2_hidden(esk4_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25, plain, (v1_xboole_0(k2_zfmisc_1(X2,X1))|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_26, negated_conjecture, (k3_zfmisc_1(esk8_0,esk9_0,esk10_0)=k3_zfmisc_1(esk11_0,esk12_0,esk13_0)&(esk8_0!=k1_xboole_0&(esk9_0!=k1_xboole_0&(esk10_0!=k1_xboole_0&(esk8_0!=esk11_0|(esk9_0!=esk12_0|esk10_0!=esk13_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).
fof(c_0_27, plain, ![X44, X45, X46]:k3_zfmisc_1(X44,X45,X46)=k2_zfmisc_1(k2_zfmisc_1(X44,X45),X46), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
cnf(c_0_28, plain, (r2_hidden(esk4_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X2)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_29, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_30, plain, (r1_tarski(k2_zfmisc_1(X1,X2),X3)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_18, c_0_25])).
fof(c_0_31, plain, (~epred2_0<=>![X3]:~v1_xboole_0(X3)), introduced(definition)).
cnf(c_0_32, negated_conjecture, (k3_zfmisc_1(esk8_0,esk9_0,esk10_0)=k3_zfmisc_1(esk11_0,esk12_0,esk13_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_34, plain, (~epred1_0<=>![X1]:~r2_hidden(X1,k1_xboole_0)), introduced(definition)).
cnf(c_0_35, plain, (~r2_hidden(X1,k2_zfmisc_1(X2,X3))|~v1_xboole_0(X3)), inference(spm,[status(thm)],[c_0_14, c_0_28])).
cnf(c_0_36, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_37, plain, (epred2_0|~v1_xboole_0(X1)), inference(split_equiv,[status(thm)],[c_0_31])).
fof(c_0_38, plain, ![X47, X48, X49]:((~r1_tarski(k2_zfmisc_1(X48,X47),k2_zfmisc_1(X49,X47))|r1_tarski(X48,X49)|X47=k1_xboole_0)&(~r1_tarski(k2_zfmisc_1(X47,X48),k2_zfmisc_1(X47,X49))|r1_tarski(X48,X49)|X47=k1_xboole_0)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t117_zfmisc_1])])])])).
cnf(c_0_39, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(esk11_0,esk12_0),esk13_0)=k2_zfmisc_1(k2_zfmisc_1(esk8_0,esk9_0),esk10_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_33]), c_0_33])).
cnf(c_0_40, plain, (~epred2_0|~epred1_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_34]), c_0_31])).
cnf(c_0_41, plain, (epred2_0), inference(spm,[status(thm)],[c_0_37, c_0_19])).
cnf(c_0_42, plain, (r1_tarski(X1,X3)|X2=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_43, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(esk8_0,esk9_0),esk10_0)=k1_xboole_0|~v1_xboole_0(esk13_0)), inference(spm,[status(thm)],[c_0_39, c_0_36])).
cnf(c_0_44, negated_conjecture, (esk10_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_45, plain, (epred1_0|~r2_hidden(X1,k1_xboole_0)), inference(split_equiv,[status(thm)],[c_0_34])).
cnf(c_0_46, plain, (~epred1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_41])])).
fof(c_0_47, plain, ![X42, X43]:((~r1_tarski(X42,k2_zfmisc_1(X42,X43))|X42=k1_xboole_0)&(~r1_tarski(X42,k2_zfmisc_1(X43,X42))|X42=k1_xboole_0)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t135_zfmisc_1])])])).
cnf(c_0_48, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk8_0,esk9_0),X1)|~v1_xboole_0(esk13_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_24])]), c_0_44])).
cnf(c_0_49, negated_conjecture, (esk9_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_50, plain, (~r2_hidden(X1,k2_zfmisc_1(X2,k1_xboole_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_28]), c_0_46])).
fof(c_0_51, plain, ![X50, X51]:(v1_xboole_0(X50)|v1_xboole_0(X51)|~v1_xboole_0(k2_zfmisc_1(X50,X51))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc10_subset_1])])])).
cnf(c_0_52, plain, (X1=k1_xboole_0|~r1_tarski(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_53, negated_conjecture, (k2_zfmisc_1(esk8_0,esk9_0)=k1_xboole_0|~v1_xboole_0(esk13_0)), inference(spm,[status(thm)],[c_0_29, c_0_48])).
cnf(c_0_54, negated_conjecture, (esk8_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_55, negated_conjecture, (r1_tarski(esk8_0,X1)|~v1_xboole_0(esk13_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_48]), c_0_49])).
fof(c_0_56, plain, ![X52, X53]:((k1_relat_1(k2_zfmisc_1(X52,X53))=X52|X53=k1_xboole_0|X52=k1_xboole_0)&(k2_relat_1(k2_zfmisc_1(X52,X53))=X53|X53=k1_xboole_0|X52=k1_xboole_0)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t195_relat_1])])])])).
cnf(c_0_57, plain, (r1_tarski(k2_zfmisc_1(X1,k1_xboole_0),X2)), inference(spm,[status(thm)],[c_0_50, c_0_15])).
cnf(c_0_58, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_59, negated_conjecture, (~v1_xboole_0(esk13_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54]), c_0_55])).
cnf(c_0_60, plain, (k2_relat_1(k2_zfmisc_1(X1,X2))=X2|X2=k1_xboole_0|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_61, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_29, c_0_57])).
cnf(c_0_62, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_36]), c_0_19])])).
cnf(c_0_63, plain, (r2_hidden(esk3_4(X1,X2,X3,X4),X1)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_64, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(esk11_0,esk12_0))|~v1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(esk8_0,esk9_0),esk10_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_39]), c_0_59])).
cnf(c_0_65, negated_conjecture, (k2_relat_1(k2_zfmisc_1(k2_zfmisc_1(esk8_0,esk9_0),esk10_0))=esk13_0|k2_zfmisc_1(esk11_0,esk12_0)=k1_xboole_0|esk13_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_60, c_0_39])).
cnf(c_0_66, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_61]), c_0_62])).
cnf(c_0_67, plain, (r2_hidden(esk1_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_68, plain, (r2_hidden(esk3_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X1)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_63])).
cnf(c_0_69, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(esk11_0,esk12_0))|~v1_xboole_0(esk10_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_36]), c_0_19])])).
cnf(c_0_70, negated_conjecture, (k2_zfmisc_1(esk11_0,esk12_0)=k1_xboole_0|k2_zfmisc_1(esk8_0,esk9_0)=k1_xboole_0|esk13_0=k1_xboole_0|esk13_0=esk10_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_65]), c_0_44])).
cnf(c_0_71, plain, (v1_xboole_0(X1)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_72, plain, (~r2_hidden(X1,k2_zfmisc_1(k1_xboole_0,X2))), inference(spm,[status(thm)],[c_0_62, c_0_68])).
cnf(c_0_73, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk11_0,esk12_0),X1)|~v1_xboole_0(esk10_0)), inference(spm,[status(thm)],[c_0_18, c_0_69])).
cnf(c_0_74, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(esk8_0,esk9_0),esk10_0)=k2_zfmisc_1(k1_xboole_0,esk13_0)|k2_zfmisc_1(esk8_0,esk9_0)=k1_xboole_0|esk13_0=esk10_0|esk13_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_39, c_0_70])).
cnf(c_0_75, negated_conjecture, (esk13_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_59, c_0_71])).
cnf(c_0_76, plain, (r1_tarski(k2_zfmisc_1(k1_xboole_0,X1),X2)), inference(spm,[status(thm)],[c_0_72, c_0_15])).
cnf(c_0_77, negated_conjecture, (k2_zfmisc_1(esk11_0,esk12_0)=k1_xboole_0|~v1_xboole_0(esk10_0)), inference(spm,[status(thm)],[c_0_29, c_0_73])).
cnf(c_0_78, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(esk8_0,esk9_0),esk10_0)=k2_zfmisc_1(k1_xboole_0,esk13_0)|k2_zfmisc_1(esk8_0,esk9_0)=k1_xboole_0|esk13_0=esk10_0), inference(sr,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_79, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_29, c_0_76])).
cnf(c_0_80, plain, (X1=k1_xboole_0|~r1_tarski(X1,k2_zfmisc_1(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_81, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(esk8_0,esk9_0),esk10_0)=k2_zfmisc_1(k1_xboole_0,esk13_0)|~v1_xboole_0(esk10_0)), inference(spm,[status(thm)],[c_0_39, c_0_77])).
cnf(c_0_82, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(esk8_0,esk9_0),esk10_0)=k1_xboole_0|k2_zfmisc_1(esk8_0,esk9_0)=k1_xboole_0|esk13_0=esk10_0), inference(rw,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_83, negated_conjecture, (~v1_xboole_0(esk10_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_44]), c_0_18])).
cnf(c_0_84, negated_conjecture, (esk13_0=esk10_0|v1_xboole_0(k2_zfmisc_1(esk8_0,esk9_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_82]), c_0_19])]), c_0_83]), c_0_71])).
cnf(c_0_85, negated_conjecture, (esk13_0=esk10_0|r1_tarski(k2_zfmisc_1(esk8_0,esk9_0),X1)), inference(spm,[status(thm)],[c_0_18, c_0_84])).
cnf(c_0_86, negated_conjecture, (esk13_0=esk10_0|r1_tarski(esk8_0,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_85]), c_0_49])).
cnf(c_0_87, negated_conjecture, (esk13_0=k1_xboole_0|r1_tarski(X1,k2_zfmisc_1(esk11_0,esk12_0))|~r1_tarski(k2_zfmisc_1(X1,esk13_0),k2_zfmisc_1(k2_zfmisc_1(esk8_0,esk9_0),esk10_0))), inference(spm,[status(thm)],[c_0_42, c_0_39])).
cnf(c_0_88, negated_conjecture, (esk13_0=esk10_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_86]), c_0_54])).
cnf(c_0_89, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_90, negated_conjecture, (esk13_0=k1_xboole_0|r1_tarski(k2_zfmisc_1(esk11_0,esk12_0),X1)|~r1_tarski(k2_zfmisc_1(k2_zfmisc_1(esk8_0,esk9_0),esk10_0),k2_zfmisc_1(X1,esk13_0))), inference(spm,[status(thm)],[c_0_42, c_0_39])).
cnf(c_0_91, negated_conjecture, (r1_tarski(X1,k2_zfmisc_1(esk11_0,esk12_0))|~r1_tarski(k2_zfmisc_1(X1,esk10_0),k2_zfmisc_1(k2_zfmisc_1(esk8_0,esk9_0),esk10_0))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_88]), c_0_88]), c_0_44])).
cnf(c_0_92, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_89])).
cnf(c_0_93, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk11_0,esk12_0),X1)|~r1_tarski(k2_zfmisc_1(k2_zfmisc_1(esk8_0,esk9_0),esk10_0),k2_zfmisc_1(X1,esk10_0))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_88]), c_0_88]), c_0_44])).
cnf(c_0_94, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk8_0,esk9_0),k2_zfmisc_1(esk11_0,esk12_0))), inference(spm,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_95, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk11_0,esk12_0),k2_zfmisc_1(esk8_0,esk9_0))), inference(spm,[status(thm)],[c_0_93, c_0_92])).
cnf(c_0_96, negated_conjecture, (k2_zfmisc_1(esk11_0,esk12_0)=k2_zfmisc_1(esk8_0,esk9_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_94]), c_0_95])])).
cnf(c_0_97, plain, (k1_relat_1(k2_zfmisc_1(X1,X2))=X1|X2=k1_xboole_0|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_98, negated_conjecture, (esk8_0!=esk11_0|esk9_0!=esk12_0|esk10_0!=esk13_0), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_99, negated_conjecture, (k2_relat_1(k2_zfmisc_1(esk8_0,esk9_0))=esk12_0|esk11_0=k1_xboole_0|esk12_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_60, c_0_96])).
cnf(c_0_100, negated_conjecture, (k1_relat_1(k2_zfmisc_1(esk8_0,esk9_0))=esk11_0|esk11_0=k1_xboole_0|esk12_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_97, c_0_96])).
cnf(c_0_101, negated_conjecture, (esk11_0!=esk8_0|esk12_0!=esk9_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_98, c_0_88])])).
cnf(c_0_102, negated_conjecture, (esk12_0=k1_xboole_0|esk11_0=k1_xboole_0|esk12_0=esk9_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_99]), c_0_49]), c_0_54])).
cnf(c_0_103, negated_conjecture, (esk12_0=k1_xboole_0|esk11_0=k1_xboole_0|esk11_0=esk8_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_100]), c_0_49]), c_0_54])).
cnf(c_0_104, negated_conjecture, (esk11_0=k1_xboole_0|esk12_0=k1_xboole_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_103])).
cnf(c_0_105, plain, (r1_tarski(X2,X3)|X1=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_106, negated_conjecture, (k2_zfmisc_1(esk8_0,esk9_0)=k1_xboole_0|esk11_0=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_104]), c_0_61])).
cnf(c_0_107, negated_conjecture, (esk11_0=k1_xboole_0|r1_tarski(esk9_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_24])]), c_0_54])).
cnf(c_0_108, negated_conjecture, (esk11_0=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_107]), c_0_49])).
cnf(c_0_109, negated_conjecture, (k2_zfmisc_1(esk8_0,esk9_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_96, c_0_108]), c_0_79])).
cnf(c_0_110, negated_conjecture, (r1_tarski(esk9_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_109]), c_0_24])]), c_0_54])).
cnf(c_0_111, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_110]), c_0_49]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 112
# Proof object clause steps            : 86
# Proof object formula steps           : 26
# Proof object conjectures             : 49
# Proof object clause conjectures      : 46
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 53
# Proof object simplifying inferences  : 57
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 30
# Processed clauses                    : 475
# ...of these trivial                  : 2
# ...subsumed                          : 194
# ...remaining for further processing  : 279
# Other redundant clauses eliminated   : 32
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 50
# Backward-rewritten                   : 100
# Generated clauses                    : 3838
# ...of the previous two non-trivial   : 3439
# Contextual simplify-reflections      : 6
# Paramodulations                      : 3762
# Factorizations                       : 17
# Equation resolutions                 : 42
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 83
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 65
# Current number of unprocessed clauses: 2953
# ...number of literals in the above   : 14268
# Current number of archived formulas  : 0
# Current number of archived clauses   : 194
# Clause-clause subsumption calls (NU) : 2085
# Rec. Clause-clause subsumption calls : 1070
# Non-unit clause-clause subsumptions  : 136
# Unit Clause-clause subsumption calls : 279
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 21
# BW rewrite match successes           : 13
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 50542

# -------------------------------------------------
# User time                : 0.041 s
# System time              : 0.003 s
# Total time               : 0.044 s
# Maximum resident set size: 28688 pages
