# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t67_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>((r1_yellow_0(X1,k7_domain_1(u1_struct_0(X2),X3,X4))&r2_hidden(k1_yellow_0(X1,k7_domain_1(u1_struct_0(X2),X3,X4)),u1_struct_0(X2)))=>(r1_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))&k1_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))=k1_yellow_0(X1,k7_domain_1(u1_struct_0(X2),X3,X4)))))))), file('data/MPTP/proofs/1236aecce7fd4ef0a087edc69a81e1af.E_input', t67_yellow_0)).
fof(t65_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>((r1_yellow_0(X1,X3)&r2_hidden(k1_yellow_0(X1,X3),u1_struct_0(X2)))=>(r1_yellow_0(X2,X3)&k1_yellow_0(X2,X3)=k1_yellow_0(X1,X3)))))), file('data/MPTP/proofs/1236aecce7fd4ef0a087edc69a81e1af.E_input', t65_yellow_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/1236aecce7fd4ef0a087edc69a81e1af.E_input', d1_xboole_0)).
fof(dt_k7_domain_1, axiom, ![X1, X2, X3]:((~(v1_xboole_0(X1))&(m1_subset_1(X2,X1)&m1_subset_1(X3,X1)))=>m1_subset_1(k7_domain_1(X1,X2,X3),k1_zfmisc_1(X1))), file('data/MPTP/proofs/1236aecce7fd4ef0a087edc69a81e1af.E_input', dt_k7_domain_1)).
fof(c_0_4, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>((r1_yellow_0(X1,k7_domain_1(u1_struct_0(X2),X3,X4))&r2_hidden(k1_yellow_0(X1,k7_domain_1(u1_struct_0(X2),X3,X4)),u1_struct_0(X2)))=>(r1_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))&k1_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))=k1_yellow_0(X1,k7_domain_1(u1_struct_0(X2),X3,X4))))))))), inference(assume_negation,[status(cth)],[t67_yellow_0])).
fof(c_0_5, plain, ![X16, X17, X18]:((r1_yellow_0(X17,X18)|(~r1_yellow_0(X16,X18)|~r2_hidden(k1_yellow_0(X16,X18),u1_struct_0(X17)))|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))|(v2_struct_0(X17)|(~v4_yellow_0(X17,X16)|~m1_yellow_0(X17,X16)))|(v2_struct_0(X16)|(~v4_orders_2(X16)|~l1_orders_2(X16))))&(k1_yellow_0(X17,X18)=k1_yellow_0(X16,X18)|(~r1_yellow_0(X16,X18)|~r2_hidden(k1_yellow_0(X16,X18),u1_struct_0(X17)))|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))|(v2_struct_0(X17)|(~v4_yellow_0(X17,X16)|~m1_yellow_0(X17,X16)))|(v2_struct_0(X16)|(~v4_orders_2(X16)|~l1_orders_2(X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t65_yellow_0])])])])])).
fof(c_0_6, negated_conjecture, ((~v2_struct_0(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0)))&((~v2_struct_0(esk2_0)&(v4_yellow_0(esk2_0,esk1_0)&m1_yellow_0(esk2_0,esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(esk2_0))&(m1_subset_1(esk4_0,u1_struct_0(esk2_0))&((r1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))&r2_hidden(k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0)),u1_struct_0(esk2_0)))&(~r1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))|k1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))!=k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_4])])])])).
fof(c_0_7, plain, ![X9, X10, X11]:((~v1_xboole_0(X9)|~r2_hidden(X10,X9))&(r2_hidden(esk5_1(X11),X11)|v1_xboole_0(X11))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_8, plain, (r1_yellow_0(X1,X2)|v2_struct_0(X1)|v2_struct_0(X3)|~r1_yellow_0(X3,X2)|~r2_hidden(k1_yellow_0(X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v4_yellow_0(X1,X3)|~m1_yellow_0(X1,X3)|~v4_orders_2(X3)|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, negated_conjecture, (r2_hidden(k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0)),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, negated_conjecture, (r1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11, negated_conjecture, (m1_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12, negated_conjecture, (v4_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_16, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_17, plain, ![X13, X14, X15]:(v1_xboole_0(X13)|(~m1_subset_1(X14,X13)|~m1_subset_1(X15,X13))|m1_subset_1(k7_domain_1(X13,X14,X15),k1_zfmisc_1(X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_domain_1])])])).
cnf(c_0_18, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_19, negated_conjecture, (r1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))|~m1_subset_1(k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10]), c_0_11]), c_0_12]), c_0_13]), c_0_14])]), c_0_15]), c_0_16])).
cnf(c_0_20, plain, (v1_xboole_0(X1)|m1_subset_1(k7_domain_1(X1,X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_22, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_23, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_18, c_0_9])).
cnf(c_0_24, negated_conjecture, (~r1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))|k1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))!=k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_25, negated_conjecture, (r1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_22])]), c_0_23])).
cnf(c_0_26, plain, (k1_yellow_0(X1,X2)=k1_yellow_0(X3,X2)|v2_struct_0(X1)|v2_struct_0(X3)|~r1_yellow_0(X3,X2)|~r2_hidden(k1_yellow_0(X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v4_yellow_0(X1,X3)|~m1_yellow_0(X1,X3)|~v4_orders_2(X3)|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_27, negated_conjecture, (k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))!=k1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25])])).
cnf(c_0_28, negated_conjecture, (~m1_subset_1(k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk4_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_9]), c_0_10]), c_0_11]), c_0_12]), c_0_13]), c_0_14])]), c_0_27]), c_0_15]), c_0_16])).
cnf(c_0_29, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_20]), c_0_21]), c_0_22])]), c_0_23]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 30
# Proof object clause steps            : 21
# Proof object formula steps           : 9
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 27
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 16
# Processed clauses                    : 37
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 37
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 6
# ...of the previous two non-trivial   : 5
# Contextual simplify-reflections      : 0
# Paramodulations                      : 6
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 19
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 5
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 18
# Clause-clause subsumption calls (NU) : 91
# Rec. Clause-clause subsumption calls : 3
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1374

# -------------------------------------------------
# User time                : 0.005 s
# System time              : 0.004 s
# Total time               : 0.009 s
# Maximum resident set size: 28684 pages
