# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C41_F1_PI_AE_Q4_CS_SP_PS_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t44_subset_1, conjecture, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(r1_xboole_0(X2,k3_subset_1(X1,X3))<=>r1_tarski(X2,X3)))), file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', t44_subset_1)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', d1_zfmisc_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', d2_subset_1)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', fc1_subset_1)).
fof(t92_xboole_1, axiom, ![X1]:k5_xboole_0(X1,X1)=k1_xboole_0, file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', t92_xboole_1)).
fof(t112_xboole_1, axiom, ![X1, X2, X3]:k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k3_xboole_0(k5_xboole_0(X1,X3),X2), file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', t112_xboole_1)).
fof(t85_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_xboole_0(X1,k4_xboole_0(X3,X2))), file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', t85_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', t100_xboole_1)).
fof(t81_xboole_1, axiom, ![X1, X2, X3]:(r1_xboole_0(X1,k4_xboole_0(X2,X3))=>r1_xboole_0(X2,k4_xboole_0(X1,X3))), file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', t81_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', commutativity_k3_xboole_0)).
fof(t5_boole, axiom, ![X1]:k5_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', t5_boole)).
fof(d5_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)), file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', d5_subset_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', t48_xboole_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', d7_xboole_0)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', t17_xboole_1)).
fof(t63_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_xboole_0(X2,X3))=>r1_xboole_0(X1,X3)), file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', t63_xboole_1)).
fof(commutativity_k5_xboole_0, axiom, ![X1, X2]:k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1), file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', commutativity_k5_xboole_0)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', idempotence_k3_xboole_0)).
fof(t47_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', t47_xboole_1)).
fof(t91_xboole_1, axiom, ![X1, X2, X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3)), file('data/MPTP/proofs/127026c1c2ce4da7810f218f62221165.E_input', t91_xboole_1)).
fof(c_0_20, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(r1_xboole_0(X2,k3_subset_1(X1,X3))<=>r1_tarski(X2,X3))))), inference(assume_negation,[status(cth)],[t44_subset_1])).
fof(c_0_21, plain, ![X24, X25, X26, X27, X28, X29]:(((~r2_hidden(X26,X25)|r1_tarski(X26,X24)|X25!=k1_zfmisc_1(X24))&(~r1_tarski(X27,X24)|r2_hidden(X27,X25)|X25!=k1_zfmisc_1(X24)))&((~r2_hidden(esk4_2(X28,X29),X29)|~r1_tarski(esk4_2(X28,X29),X28)|X29=k1_zfmisc_1(X28))&(r2_hidden(esk4_2(X28,X29),X29)|r1_tarski(esk4_2(X28,X29),X28)|X29=k1_zfmisc_1(X28)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
fof(c_0_22, plain, ![X49, X50]:(((~m1_subset_1(X50,X49)|r2_hidden(X50,X49)|v1_xboole_0(X49))&(~r2_hidden(X50,X49)|m1_subset_1(X50,X49)|v1_xboole_0(X49)))&((~m1_subset_1(X50,X49)|v1_xboole_0(X50)|~v1_xboole_0(X49))&(~v1_xboole_0(X50)|m1_subset_1(X50,X49)|~v1_xboole_0(X49)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_23, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))&((~r1_xboole_0(esk2_0,k3_subset_1(esk1_0,esk3_0))|~r1_tarski(esk2_0,esk3_0))&(r1_xboole_0(esk2_0,k3_subset_1(esk1_0,esk3_0))|r1_tarski(esk2_0,esk3_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).
fof(c_0_24, plain, ![X17]:~v1_xboole_0(k1_zfmisc_1(X17)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
fof(c_0_25, plain, ![X7]:k5_xboole_0(X7,X7)=k1_xboole_0, inference(variable_rename,[status(thm)],[t92_xboole_1])).
fof(c_0_26, plain, ![X9, X10, X11]:k5_xboole_0(k3_xboole_0(X9,X10),k3_xboole_0(X11,X10))=k3_xboole_0(k5_xboole_0(X9,X11),X10), inference(variable_rename,[status(thm)],[t112_xboole_1])).
fof(c_0_27, plain, ![X36, X37, X38]:(~r1_tarski(X36,X37)|r1_xboole_0(X36,k4_xboole_0(X38,X37))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t85_xboole_1])])).
fof(c_0_28, plain, ![X42, X43]:k4_xboole_0(X42,X43)=k5_xboole_0(X42,k3_xboole_0(X42,X43)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_29, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_33, plain, ![X44, X45, X46]:(~r1_xboole_0(X44,k4_xboole_0(X45,X46))|r1_xboole_0(X45,k4_xboole_0(X44,X46))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t81_xboole_1])])).
fof(c_0_34, plain, ![X20, X21]:k3_xboole_0(X20,X21)=k3_xboole_0(X21,X20), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_35, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_36, plain, (k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k3_xboole_0(k5_xboole_0(X1,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_37, plain, (r1_xboole_0(X1,k4_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_38, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_39, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(er,[status(thm)],[c_0_29])).
cnf(c_0_40, negated_conjecture, (r2_hidden(esk3_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])).
cnf(c_0_41, plain, (r1_xboole_0(X2,k4_xboole_0(X1,X3))|~r1_xboole_0(X1,k4_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_35])).
fof(c_0_44, plain, ![X16]:k5_xboole_0(X16,k1_xboole_0)=X16, inference(variable_rename,[status(thm)],[t5_boole])).
fof(c_0_45, plain, ![X31, X32]:(~m1_subset_1(X32,k1_zfmisc_1(X31))|k3_subset_1(X31,X32)=k4_xboole_0(X31,X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
cnf(c_0_46, plain, (r1_xboole_0(X1,k5_xboole_0(X3,k3_xboole_0(X3,X2)))|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_47, negated_conjecture, (r1_tarski(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_48, plain, (r1_xboole_0(X2,k5_xboole_0(X1,k3_xboole_0(X1,X3)))|~r1_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_38]), c_0_38])).
cnf(c_0_49, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_50, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_51, plain, (k3_subset_1(X2,X1)=k4_xboole_0(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_52, plain, ![X14, X15]:k4_xboole_0(X14,k4_xboole_0(X14,X15))=k3_xboole_0(X14,X15), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_53, plain, ![X22, X23]:((~r1_xboole_0(X22,X23)|k3_xboole_0(X22,X23)=k1_xboole_0)&(k3_xboole_0(X22,X23)!=k1_xboole_0|r1_xboole_0(X22,X23))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_54, negated_conjecture, (r1_xboole_0(esk3_0,k5_xboole_0(X1,k3_xboole_0(X1,esk1_0)))), inference(spm,[status(thm)],[c_0_46, c_0_47])).
fof(c_0_55, plain, ![X18, X19]:r1_tarski(k3_xboole_0(X18,X19),X18), inference(variable_rename,[status(thm)],[t17_xboole_1])).
cnf(c_0_56, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_49]), c_0_50]), c_0_50])).
cnf(c_0_57, negated_conjecture, (r1_xboole_0(esk2_0,k3_subset_1(esk1_0,esk3_0))|r1_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_58, plain, (k5_xboole_0(X2,k3_xboole_0(X2,X1))=k3_subset_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_51, c_0_38])).
cnf(c_0_59, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_60, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_61, negated_conjecture, (r1_xboole_0(X1,k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,esk1_0)))), inference(spm,[status(thm)],[c_0_48, c_0_54])).
fof(c_0_62, plain, ![X33, X34, X35]:(~r1_tarski(X33,X34)|~r1_xboole_0(X34,X35)|r1_xboole_0(X33,X35)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_xboole_1])])).
cnf(c_0_63, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_64, negated_conjecture, (r1_tarski(esk2_0,esk3_0)|r1_xboole_0(k3_subset_1(esk1_0,esk3_0),esk2_0)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_65, negated_conjecture, (k3_subset_1(esk1_0,esk3_0)=k5_xboole_0(esk1_0,k3_xboole_0(esk3_0,esk1_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_31]), c_0_42])).
cnf(c_0_66, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2))))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_38]), c_0_38])).
cnf(c_0_67, negated_conjecture, (k3_xboole_0(X1,k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,esk1_0)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_60, c_0_61])).
fof(c_0_68, plain, ![X12, X13]:k5_xboole_0(X12,X13)=k5_xboole_0(X13,X12), inference(variable_rename,[status(thm)],[commutativity_k5_xboole_0])).
fof(c_0_69, plain, ![X8]:k3_xboole_0(X8,X8)=X8, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_70, plain, (r1_xboole_0(X1,X3)|~r1_tarski(X1,X2)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_71, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_63, c_0_42])).
cnf(c_0_72, negated_conjecture, (r1_tarski(esk2_0,esk3_0)|r1_xboole_0(k5_xboole_0(esk1_0,k3_xboole_0(esk3_0,esk1_0)),esk2_0)), inference(rw,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_73, negated_conjecture, (k3_xboole_0(esk3_0,esk1_0)=esk3_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_50])).
cnf(c_0_74, plain, (k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_75, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_76, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_77, plain, (r1_xboole_0(k3_xboole_0(X1,X2),X3)|~r1_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_78, negated_conjecture, (r1_tarski(esk2_0,esk3_0)|r1_xboole_0(k5_xboole_0(esk3_0,esk1_0),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_73]), c_0_74])).
cnf(c_0_79, negated_conjecture, (r2_hidden(esk2_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_75]), c_0_32])).
cnf(c_0_80, plain, (k5_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X1,k5_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_76]), c_0_74]), c_0_42])).
cnf(c_0_81, negated_conjecture, (r1_tarski(esk2_0,esk3_0)|r1_xboole_0(k3_xboole_0(X1,k5_xboole_0(esk3_0,esk1_0)),esk2_0)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
fof(c_0_82, plain, ![X47, X48]:k4_xboole_0(X47,k3_xboole_0(X47,X48))=k4_xboole_0(X47,X48), inference(variable_rename,[status(thm)],[t47_xboole_1])).
fof(c_0_83, plain, ![X39, X40, X41]:k5_xboole_0(k5_xboole_0(X39,X40),X41)=k5_xboole_0(X39,k5_xboole_0(X40,X41)), inference(variable_rename,[status(thm)],[t91_xboole_1])).
cnf(c_0_84, negated_conjecture, (r1_tarski(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_39, c_0_79])).
cnf(c_0_85, plain, (r1_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|~r1_xboole_0(X2,k3_xboole_0(X1,k5_xboole_0(X3,X1)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_42]), c_0_80])).
cnf(c_0_86, negated_conjecture, (r1_tarski(esk2_0,esk3_0)|r1_xboole_0(esk2_0,k3_xboole_0(X1,k5_xboole_0(esk3_0,esk1_0)))), inference(spm,[status(thm)],[c_0_56, c_0_81])).
cnf(c_0_87, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_88, plain, (k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_89, negated_conjecture, (r1_xboole_0(esk2_0,X1)|~r1_xboole_0(esk1_0,X1)), inference(spm,[status(thm)],[c_0_70, c_0_84])).
cnf(c_0_90, negated_conjecture, (r1_tarski(esk2_0,esk3_0)|r1_xboole_0(esk1_0,k5_xboole_0(esk2_0,k3_xboole_0(esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_91, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k3_xboole_0(X1,X2)))=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_38]), c_0_38])).
cnf(c_0_92, plain, (k5_xboole_0(X1,k5_xboole_0(X2,X3))=k5_xboole_0(X3,k5_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_74, c_0_88])).
cnf(c_0_93, plain, (k5_xboole_0(k1_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_50, c_0_74])).
cnf(c_0_94, negated_conjecture, (r1_xboole_0(esk2_0,k5_xboole_0(esk2_0,k3_xboole_0(esk2_0,esk3_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_46])).
cnf(c_0_95, plain, (k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2)))=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_66]), c_0_91])).
cnf(c_0_96, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k3_xboole_0(X2,X1)))=k3_xboole_0(X1,k5_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_42]), c_0_80])).
cnf(c_0_97, negated_conjecture, (k3_xboole_0(esk1_0,k5_xboole_0(esk3_0,esk1_0))=k5_xboole_0(esk3_0,esk1_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_73]), c_0_74])).
cnf(c_0_98, plain, (k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X1,X2)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_35, c_0_88])).
cnf(c_0_99, plain, (k5_xboole_0(X1,k5_xboole_0(X2,X1))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_35]), c_0_50])).
cnf(c_0_100, negated_conjecture, (~r1_xboole_0(esk2_0,k3_subset_1(esk1_0,esk3_0))|~r1_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_101, plain, (k5_xboole_0(X1,k5_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_35]), c_0_93])).
cnf(c_0_102, negated_conjecture, (k5_xboole_0(esk2_0,k3_xboole_0(esk2_0,esk3_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_94]), c_0_95])).
cnf(c_0_103, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_104, negated_conjecture, (k3_xboole_0(esk3_0,k5_xboole_0(esk3_0,esk1_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_76]), c_0_88]), c_0_98]), c_0_99]), c_0_42])).
cnf(c_0_105, negated_conjecture, (~r1_tarski(esk2_0,esk3_0)|~r1_xboole_0(esk2_0,k5_xboole_0(esk1_0,k3_xboole_0(esk3_0,esk1_0)))), inference(rw,[status(thm)],[c_0_100, c_0_65])).
cnf(c_0_106, negated_conjecture, (k3_xboole_0(esk2_0,esk3_0)=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_50])).
cnf(c_0_107, negated_conjecture, (r1_xboole_0(esk3_0,k5_xboole_0(esk3_0,esk1_0))), inference(spm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_108, negated_conjecture, (~r1_tarski(esk2_0,esk3_0)|~r1_xboole_0(esk2_0,k5_xboole_0(esk3_0,esk1_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_105, c_0_73]), c_0_74])).
cnf(c_0_109, negated_conjecture, (r1_tarski(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_71, c_0_106])).
cnf(c_0_110, negated_conjecture, (r1_xboole_0(k3_xboole_0(X1,esk3_0),k5_xboole_0(esk3_0,esk1_0))), inference(spm,[status(thm)],[c_0_77, c_0_107])).
cnf(c_0_111, negated_conjecture, (~r1_xboole_0(esk2_0,k5_xboole_0(esk3_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_108, c_0_109])])).
cnf(c_0_112, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_106]), c_0_111]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 113
# Proof object clause steps            : 72
# Proof object formula steps           : 41
# Proof object conjectures             : 35
# Proof object clause conjectures      : 32
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 38
# Proof object simplifying inferences  : 41
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 20
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 29
# Processed clauses                    : 398
# ...of these trivial                  : 85
# ...subsumed                          : 81
# ...remaining for further processing  : 231
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 46
# Generated clauses                    : 2261
# ...of the previous two non-trivial   : 1375
# Contextual simplify-reflections      : 1
# Paramodulations                      : 2258
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 156
#    Positive orientable unit clauses  : 106
#    Positive unorientable unit clauses: 4
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 44
# Current number of unprocessed clauses: 1009
# ...number of literals in the above   : 1317
# Current number of archived formulas  : 0
# Current number of archived clauses   : 76
# Clause-clause subsumption calls (NU) : 801
# Rec. Clause-clause subsumption calls : 776
# Non-unit clause-clause subsumptions  : 55
# Unit Clause-clause subsumption calls : 493
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 146
# BW rewrite match successes           : 48
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 24530

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.006 s
# Total time               : 0.023 s
# Maximum resident set size: 28848 pages
