# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/1279f4566ea24b1589a517bcb29f7acf.E_input', d5_relat_1)).
fof(d11_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(v1_relat_1(X3)=>(X3=k7_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>(r2_hidden(X4,X2)&r2_hidden(k4_tarski(X4,X5),X1)))))), file('data/MPTP/proofs/1279f4566ea24b1589a517bcb29f7acf.E_input', d11_relat_1)).
fof(dt_k7_relat_1, axiom, ![X1, X2]:(v1_relat_1(X1)=>v1_relat_1(k7_relat_1(X1,X2))), file('data/MPTP/proofs/1279f4566ea24b1589a517bcb29f7acf.E_input', dt_k7_relat_1)).
fof(t8_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>(r2_hidden(X1,k1_relat_1(X3))&X2=k1_funct_1(X3,X1)))), file('data/MPTP/proofs/1279f4566ea24b1589a517bcb29f7acf.E_input', t8_funct_1)).
fof(t73_funct_1, conjecture, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((r2_hidden(X1,k1_relat_1(X3))&r2_hidden(X1,X2))=>r2_hidden(k1_funct_1(X3,X1),k2_relat_1(k7_relat_1(X3,X2))))), file('data/MPTP/proofs/1279f4566ea24b1589a517bcb29f7acf.E_input', t73_funct_1)).
fof(c_0_5, plain, ![X6, X7, X8, X10, X11, X12, X13, X15]:(((~r2_hidden(X8,X7)|r2_hidden(k4_tarski(esk1_3(X6,X7,X8),X8),X6)|X7!=k2_relat_1(X6))&(~r2_hidden(k4_tarski(X11,X10),X6)|r2_hidden(X10,X7)|X7!=k2_relat_1(X6)))&((~r2_hidden(esk2_2(X12,X13),X13)|~r2_hidden(k4_tarski(X15,esk2_2(X12,X13)),X12)|X13=k2_relat_1(X12))&(r2_hidden(esk2_2(X12,X13),X13)|r2_hidden(k4_tarski(esk3_2(X12,X13),esk2_2(X12,X13)),X12)|X13=k2_relat_1(X12)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
fof(c_0_6, plain, ![X17, X18, X19, X20, X21, X22, X23]:((((r2_hidden(X20,X18)|~r2_hidden(k4_tarski(X20,X21),X19)|X19!=k7_relat_1(X17,X18)|~v1_relat_1(X19)|~v1_relat_1(X17))&(r2_hidden(k4_tarski(X20,X21),X17)|~r2_hidden(k4_tarski(X20,X21),X19)|X19!=k7_relat_1(X17,X18)|~v1_relat_1(X19)|~v1_relat_1(X17)))&(~r2_hidden(X22,X18)|~r2_hidden(k4_tarski(X22,X23),X17)|r2_hidden(k4_tarski(X22,X23),X19)|X19!=k7_relat_1(X17,X18)|~v1_relat_1(X19)|~v1_relat_1(X17)))&((~r2_hidden(k4_tarski(esk4_3(X17,X18,X19),esk5_3(X17,X18,X19)),X19)|(~r2_hidden(esk4_3(X17,X18,X19),X18)|~r2_hidden(k4_tarski(esk4_3(X17,X18,X19),esk5_3(X17,X18,X19)),X17))|X19=k7_relat_1(X17,X18)|~v1_relat_1(X19)|~v1_relat_1(X17))&((r2_hidden(esk4_3(X17,X18,X19),X18)|r2_hidden(k4_tarski(esk4_3(X17,X18,X19),esk5_3(X17,X18,X19)),X19)|X19=k7_relat_1(X17,X18)|~v1_relat_1(X19)|~v1_relat_1(X17))&(r2_hidden(k4_tarski(esk4_3(X17,X18,X19),esk5_3(X17,X18,X19)),X17)|r2_hidden(k4_tarski(esk4_3(X17,X18,X19),esk5_3(X17,X18,X19)),X19)|X19=k7_relat_1(X17,X18)|~v1_relat_1(X19)|~v1_relat_1(X17))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_relat_1])])])])])])).
fof(c_0_7, plain, ![X26, X27]:(~v1_relat_1(X26)|v1_relat_1(k7_relat_1(X26,X27))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_relat_1])])).
cnf(c_0_8, plain, (r2_hidden(X2,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (r2_hidden(k4_tarski(X1,X3),X5)|~r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),X4)|X5!=k7_relat_1(X4,X2)|~v1_relat_1(X5)|~v1_relat_1(X4)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_11, plain, ![X28, X29, X30]:(((r2_hidden(X28,k1_relat_1(X30))|~r2_hidden(k4_tarski(X28,X29),X30)|(~v1_relat_1(X30)|~v1_funct_1(X30)))&(X29=k1_funct_1(X30,X28)|~r2_hidden(k4_tarski(X28,X29),X30)|(~v1_relat_1(X30)|~v1_funct_1(X30))))&(~r2_hidden(X28,k1_relat_1(X30))|X29!=k1_funct_1(X30,X28)|r2_hidden(k4_tarski(X28,X29),X30)|(~v1_relat_1(X30)|~v1_funct_1(X30)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_funct_1])])])).
fof(c_0_12, negated_conjecture, ~(![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((r2_hidden(X1,k1_relat_1(X3))&r2_hidden(X1,X2))=>r2_hidden(k1_funct_1(X3,X1),k2_relat_1(k7_relat_1(X3,X2)))))), inference(assume_negation,[status(cth)],[t73_funct_1])).
cnf(c_0_13, plain, (r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)), inference(er,[status(thm)],[c_0_8])).
cnf(c_0_14, plain, (r2_hidden(k4_tarski(X1,X2),k7_relat_1(X3,X4))|~v1_relat_1(X3)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,X4)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_9]), c_0_10])).
cnf(c_0_15, plain, (r2_hidden(k4_tarski(X1,X3),X2)|~r2_hidden(X1,k1_relat_1(X2))|X3!=k1_funct_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_16, negated_conjecture, ((v1_relat_1(esk8_0)&v1_funct_1(esk8_0))&((r2_hidden(esk6_0,k1_relat_1(esk8_0))&r2_hidden(esk6_0,esk7_0))&~r2_hidden(k1_funct_1(esk8_0,esk6_0),k2_relat_1(k7_relat_1(esk8_0,esk7_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
cnf(c_0_17, plain, (r2_hidden(X1,k2_relat_1(k7_relat_1(X2,X3)))|~v1_relat_1(X2)|~r2_hidden(k4_tarski(X4,X1),X2)|~r2_hidden(X4,X3)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_18, plain, (r2_hidden(k4_tarski(X1,k1_funct_1(X2,X1)),X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_15])).
cnf(c_0_19, negated_conjecture, (~r2_hidden(k1_funct_1(esk8_0,esk6_0),k2_relat_1(k7_relat_1(esk8_0,esk7_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relat_1(k7_relat_1(X1,X3)))|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))|~r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_21, negated_conjecture, (v1_funct_1(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, negated_conjecture, (v1_relat_1(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, negated_conjecture, (r2_hidden(esk6_0,k1_relat_1(esk8_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, negated_conjecture, (r2_hidden(esk6_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 26
# Proof object clause steps            : 15
# Proof object formula steps           : 11
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 9
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 19
# Processed clauses                    : 76
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 75
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 224
# ...of the previous two non-trivial   : 210
# Contextual simplify-reflections      : 7
# Paramodulations                      : 216
# Factorizations                       : 2
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 50
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 45
# Current number of unprocessed clauses: 171
# ...number of literals in the above   : 1005
# Current number of archived formulas  : 0
# Current number of archived clauses   : 19
# Clause-clause subsumption calls (NU) : 890
# Rec. Clause-clause subsumption calls : 136
# Non-unit clause-clause subsumptions  : 8
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7964

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.000 s
# Total time               : 0.013 s
# Maximum resident set size: 28760 pages
