# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C47_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/128eb6ff400c4ee998ac8539fcf5a5df.E_input', t12_setfam_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/128eb6ff400c4ee998ac8539fcf5a5df.E_input', t70_enumset1)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/128eb6ff400c4ee998ac8539fcf5a5df.E_input', d4_xboole_0)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/128eb6ff400c4ee998ac8539fcf5a5df.E_input', t71_enumset1)).
fof(t182_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>k1_relat_1(k5_relat_1(X1,X2))=k10_relat_1(X1,k1_relat_1(X2)))), file('data/MPTP/proofs/128eb6ff400c4ee998ac8539fcf5a5df.E_input', t182_relat_1)).
fof(t71_relat_1, axiom, ![X1]:(k1_relat_1(k6_relat_1(X1))=X1&k2_relat_1(k6_relat_1(X1))=X1), file('data/MPTP/proofs/128eb6ff400c4ee998ac8539fcf5a5df.E_input', t71_relat_1)).
fof(fc3_funct_1, axiom, ![X1]:(v1_relat_1(k6_relat_1(X1))&v1_funct_1(k6_relat_1(X1))), file('data/MPTP/proofs/128eb6ff400c4ee998ac8539fcf5a5df.E_input', fc3_funct_1)).
fof(t37_funct_1, conjecture, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>k1_relat_1(k5_relat_1(k6_relat_1(X1),X2))=k3_xboole_0(k1_relat_1(X2),X1)), file('data/MPTP/proofs/128eb6ff400c4ee998ac8539fcf5a5df.E_input', t37_funct_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/128eb6ff400c4ee998ac8539fcf5a5df.E_input', t28_xboole_1)).
fof(t94_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k7_relat_1(X2,X1)=k5_relat_1(k6_relat_1(X1),X2)), file('data/MPTP/proofs/128eb6ff400c4ee998ac8539fcf5a5df.E_input', t94_relat_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/128eb6ff400c4ee998ac8539fcf5a5df.E_input', commutativity_k2_tarski)).
fof(t21_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k1_relat_1(k5_relat_1(X3,X2)))<=>(r2_hidden(X1,k1_relat_1(X3))&r2_hidden(k1_funct_1(X3,X1),k1_relat_1(X2)))))), file('data/MPTP/proofs/128eb6ff400c4ee998ac8539fcf5a5df.E_input', t21_funct_1)).
fof(t87_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k1_relat_1(k7_relat_1(X2,X1)),X1)), file('data/MPTP/proofs/128eb6ff400c4ee998ac8539fcf5a5df.E_input', t87_relat_1)).
fof(t35_funct_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>k1_funct_1(k6_relat_1(X2),X1)=X1), file('data/MPTP/proofs/128eb6ff400c4ee998ac8539fcf5a5df.E_input', t35_funct_1)).
fof(c_0_14, plain, ![X7, X8]:k1_setfam_1(k2_tarski(X7,X8))=k3_xboole_0(X7,X8), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_15, plain, ![X20, X21]:k1_enumset1(X20,X20,X21)=k2_tarski(X20,X21), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_16, plain, ![X31, X32, X33, X34, X35, X36, X37, X38]:((((r2_hidden(X34,X31)|~r2_hidden(X34,X33)|X33!=k3_xboole_0(X31,X32))&(r2_hidden(X34,X32)|~r2_hidden(X34,X33)|X33!=k3_xboole_0(X31,X32)))&(~r2_hidden(X35,X31)|~r2_hidden(X35,X32)|r2_hidden(X35,X33)|X33!=k3_xboole_0(X31,X32)))&((~r2_hidden(esk3_3(X36,X37,X38),X38)|(~r2_hidden(esk3_3(X36,X37,X38),X36)|~r2_hidden(esk3_3(X36,X37,X38),X37))|X38=k3_xboole_0(X36,X37))&((r2_hidden(esk3_3(X36,X37,X38),X36)|r2_hidden(esk3_3(X36,X37,X38),X38)|X38=k3_xboole_0(X36,X37))&(r2_hidden(esk3_3(X36,X37,X38),X37)|r2_hidden(esk3_3(X36,X37,X38),X38)|X38=k3_xboole_0(X36,X37))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_17, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_19, plain, ![X13, X14, X15]:k2_enumset1(X13,X13,X14,X15)=k1_enumset1(X13,X14,X15), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_20, plain, ![X27, X28]:(~v1_relat_1(X27)|(~v1_relat_1(X28)|k1_relat_1(k5_relat_1(X27,X28))=k10_relat_1(X27,k1_relat_1(X28)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t182_relat_1])])])).
fof(c_0_21, plain, ![X16]:(k1_relat_1(k6_relat_1(X16))=X16&k2_relat_1(k6_relat_1(X16))=X16), inference(variable_rename,[status(thm)],[t71_relat_1])).
fof(c_0_22, plain, ![X24]:(v1_relat_1(k6_relat_1(X24))&v1_funct_1(k6_relat_1(X24))), inference(variable_rename,[status(thm)],[fc3_funct_1])).
fof(c_0_23, negated_conjecture, ~(![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>k1_relat_1(k5_relat_1(k6_relat_1(X1),X2))=k3_xboole_0(k1_relat_1(X2),X1))), inference(assume_negation,[status(cth)],[t37_funct_1])).
cnf(c_0_24, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_26, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_27, plain, ![X11, X12]:(~r1_tarski(X11,X12)|k3_xboole_0(X11,X12)=X11), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_28, plain, (k1_relat_1(k5_relat_1(X1,X2))=k10_relat_1(X1,k1_relat_1(X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, plain, (k1_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, plain, (v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_31, plain, ![X9, X10]:(~v1_relat_1(X10)|k7_relat_1(X10,X9)=k5_relat_1(k6_relat_1(X9),X10)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t94_relat_1])])).
fof(c_0_32, negated_conjecture, ((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0))!=k3_xboole_0(k1_relat_1(esk2_0),esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])).
fof(c_0_33, plain, ![X22, X23]:k2_tarski(X22,X23)=k2_tarski(X23,X22), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_34, plain, (r2_hidden(esk3_3(X1,X2,X3),X1)|r2_hidden(esk3_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_35, plain, (r2_hidden(esk3_3(X1,X2,X3),X2)|r2_hidden(esk3_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_36, plain, ![X17, X18, X19]:(((r2_hidden(X17,k1_relat_1(X19))|~r2_hidden(X17,k1_relat_1(k5_relat_1(X19,X18)))|(~v1_relat_1(X19)|~v1_funct_1(X19))|(~v1_relat_1(X18)|~v1_funct_1(X18)))&(r2_hidden(k1_funct_1(X19,X17),k1_relat_1(X18))|~r2_hidden(X17,k1_relat_1(k5_relat_1(X19,X18)))|(~v1_relat_1(X19)|~v1_funct_1(X19))|(~v1_relat_1(X18)|~v1_funct_1(X18))))&(~r2_hidden(X17,k1_relat_1(X19))|~r2_hidden(k1_funct_1(X19,X17),k1_relat_1(X18))|r2_hidden(X17,k1_relat_1(k5_relat_1(X19,X18)))|(~v1_relat_1(X19)|~v1_funct_1(X19))|(~v1_relat_1(X18)|~v1_funct_1(X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t21_funct_1])])])])).
cnf(c_0_37, plain, (r2_hidden(X1,X2)|X3!=k1_setfam_1(k2_enumset1(X4,X4,X4,X2))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25]), c_0_26])).
cnf(c_0_38, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_39, plain, ![X29, X30]:(~v1_relat_1(X30)|r1_tarski(k1_relat_1(k7_relat_1(X30,X29)),X29)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t87_relat_1])])).
cnf(c_0_40, plain, (k1_relat_1(k5_relat_1(X1,k6_relat_1(X2)))=k10_relat_1(X1,X2)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])])).
cnf(c_0_41, plain, (k7_relat_1(X1,X2)=k5_relat_1(k6_relat_1(X2),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_42, negated_conjecture, (k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0))!=k3_xboole_0(k1_relat_1(esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_44, plain, (X3=k3_xboole_0(X1,X2)|~r2_hidden(esk3_3(X1,X2,X3),X3)|~r2_hidden(esk3_3(X1,X2,X3),X1)|~r2_hidden(esk3_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_45, plain, (X3=k1_setfam_1(k2_enumset1(X1,X1,X1,X2))|r2_hidden(esk3_3(X1,X2,X3),X3)|r2_hidden(esk3_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_25]), c_0_26])).
cnf(c_0_46, plain, (X3=k1_setfam_1(k2_enumset1(X1,X1,X1,X2))|r2_hidden(esk3_3(X1,X2,X3),X3)|r2_hidden(esk3_3(X1,X2,X3),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_25]), c_0_26])).
fof(c_0_47, plain, ![X25, X26]:(~r2_hidden(X25,X26)|k1_funct_1(k6_relat_1(X26),X25)=X25), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t35_funct_1])])).
cnf(c_0_48, plain, (r2_hidden(k1_funct_1(X1,X2),k1_relat_1(X3))|~r2_hidden(X2,k1_relat_1(k5_relat_1(X1,X3)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_49, plain, (v1_funct_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_50, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_setfam_1(k2_enumset1(X3,X3,X3,X2)))), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_51, plain, (k1_setfam_1(k2_enumset1(X1,X1,X1,X2))=X1|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_25]), c_0_26])).
cnf(c_0_52, plain, (r1_tarski(k1_relat_1(k7_relat_1(X1,X2)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_53, plain, (k1_relat_1(k7_relat_1(k6_relat_1(X1),X2))=k10_relat_1(k6_relat_1(X2),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_30]), c_0_30])])).
cnf(c_0_54, negated_conjecture, (k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0))!=k1_setfam_1(k2_enumset1(k1_relat_1(esk2_0),k1_relat_1(esk2_0),k1_relat_1(esk2_0),esk1_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_25]), c_0_26])).
cnf(c_0_55, plain, (k2_enumset1(X1,X1,X1,X2)=k2_enumset1(X2,X2,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_18]), c_0_18]), c_0_26]), c_0_26])).
cnf(c_0_56, plain, (X3=k1_setfam_1(k2_enumset1(X1,X1,X1,X2))|~r2_hidden(esk3_3(X1,X2,X3),X3)|~r2_hidden(esk3_3(X1,X2,X3),X2)|~r2_hidden(esk3_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_25]), c_0_26])).
cnf(c_0_57, plain, (k1_setfam_1(k2_enumset1(X1,X1,X1,X2))=X1|r2_hidden(esk3_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_45])).
cnf(c_0_58, plain, (k1_setfam_1(k2_enumset1(X1,X1,X1,X2))=X2|r2_hidden(esk3_3(X1,X2,X2),X2)), inference(ef,[status(thm)],[c_0_46])).
cnf(c_0_59, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X4!=k3_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_60, plain, (r2_hidden(X1,k1_relat_1(k5_relat_1(X2,X3)))|~r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k1_funct_1(X2,X1),k1_relat_1(X3))|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_61, plain, (k1_funct_1(k6_relat_1(X2),X1)=X1|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_62, plain, (r2_hidden(k1_funct_1(X1,X2),X3)|~r2_hidden(X2,k1_relat_1(k5_relat_1(X1,k6_relat_1(X3))))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_29]), c_0_49]), c_0_30])])).
cnf(c_0_63, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_64, plain, (r1_tarski(k10_relat_1(k6_relat_1(X1),X2),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_30])])).
cnf(c_0_65, negated_conjecture, (k1_setfam_1(k2_enumset1(esk1_0,esk1_0,esk1_0,k1_relat_1(esk2_0)))!=k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0))), inference(rw,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_66, plain, (k1_setfam_1(k2_enumset1(X1,X1,X1,X2))=X1|~r2_hidden(esk3_3(X1,X2,X1),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_57])).
cnf(c_0_67, plain, (k1_setfam_1(k2_enumset1(X1,X1,X1,X2))=X1|r2_hidden(esk3_3(X2,X1,X1),X1)), inference(spm,[status(thm)],[c_0_58, c_0_55])).
cnf(c_0_68, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_69, plain, (r2_hidden(X1,X4)|X4!=k1_setfam_1(k2_enumset1(X2,X2,X2,X3))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_25]), c_0_26])).
cnf(c_0_70, plain, (r2_hidden(X1,k1_relat_1(k5_relat_1(k6_relat_1(X2),X3)))|~r2_hidden(X1,k1_relat_1(X3))|~r2_hidden(X1,X2)|~v1_funct_1(X3)|~v1_relat_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_29]), c_0_49]), c_0_30])])).
cnf(c_0_71, plain, (r2_hidden(k1_funct_1(k6_relat_1(X1),X2),X3)|~r2_hidden(X2,k10_relat_1(k6_relat_1(X1),X3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_41]), c_0_53]), c_0_49]), c_0_30]), c_0_30])])).
cnf(c_0_72, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k10_relat_1(k6_relat_1(X2),X3))), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_73, negated_conjecture, (r2_hidden(esk3_3(X1,X2,k1_setfam_1(k2_enumset1(esk1_0,esk1_0,esk1_0,k1_relat_1(esk2_0)))),k1_setfam_1(k2_enumset1(esk1_0,esk1_0,esk1_0,k1_relat_1(esk2_0))))|r2_hidden(esk3_3(X1,X2,k1_setfam_1(k2_enumset1(esk1_0,esk1_0,esk1_0,k1_relat_1(esk2_0)))),X2)|k1_setfam_1(k2_enumset1(X1,X1,X1,X2))!=k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0))), inference(spm,[status(thm)],[c_0_65, c_0_46])).
cnf(c_0_74, plain, (k1_setfam_1(k2_enumset1(X1,X1,X1,X1))=X1), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_75, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(X1,k1_relat_1(k5_relat_1(X2,X3)))|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_76, plain, (r2_hidden(X1,X2)|X3!=k1_setfam_1(k2_enumset1(X2,X2,X2,X4))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_25]), c_0_26])).
cnf(c_0_77, plain, (r2_hidden(X1,k1_setfam_1(k2_enumset1(X2,X2,X2,X3)))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_69])).
cnf(c_0_78, plain, (r2_hidden(X1,k10_relat_1(k6_relat_1(X2),X3))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_40]), c_0_29]), c_0_49]), c_0_30]), c_0_30])])).
cnf(c_0_79, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k10_relat_1(k6_relat_1(X3),X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_61]), c_0_72])).
cnf(c_0_80, negated_conjecture, (r2_hidden(esk3_3(k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)),k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)),k1_setfam_1(k2_enumset1(esk1_0,esk1_0,esk1_0,k1_relat_1(esk2_0)))),k1_setfam_1(k2_enumset1(esk1_0,esk1_0,esk1_0,k1_relat_1(esk2_0))))|r2_hidden(esk3_3(k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)),k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)),k1_setfam_1(k2_enumset1(esk1_0,esk1_0,esk1_0,k1_relat_1(esk2_0)))),k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)))), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74])])).
cnf(c_0_81, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(X1,k10_relat_1(X2,X3))|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_40]), c_0_49]), c_0_30])])).
cnf(c_0_82, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_setfam_1(k2_enumset1(X2,X2,X2,X3)))), inference(er,[status(thm)],[c_0_76])).
cnf(c_0_83, plain, (k1_setfam_1(k2_enumset1(X1,X1,X1,X2))=k1_setfam_1(k2_enumset1(X3,X3,X3,X4))|~r2_hidden(esk3_3(X3,X4,k1_setfam_1(k2_enumset1(X1,X1,X1,X2))),X4)|~r2_hidden(esk3_3(X3,X4,k1_setfam_1(k2_enumset1(X1,X1,X1,X2))),X3)|~r2_hidden(esk3_3(X3,X4,k1_setfam_1(k2_enumset1(X1,X1,X1,X2))),X2)|~r2_hidden(esk3_3(X3,X4,k1_setfam_1(k2_enumset1(X1,X1,X1,X2))),X1)), inference(spm,[status(thm)],[c_0_56, c_0_77])).
cnf(c_0_84, plain, (r2_hidden(X1,k1_relat_1(k5_relat_1(k6_relat_1(X2),X3)))|~r2_hidden(X1,k1_relat_1(X3))|~r2_hidden(X1,X2)|~v1_relat_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_28]), c_0_30])])).
cnf(c_0_85, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(X1,k1_relat_1(k5_relat_1(k6_relat_1(X3),X2)))|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_28]), c_0_30])])).
cnf(c_0_86, negated_conjecture, (r2_hidden(esk3_3(k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)),k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)),k1_setfam_1(k2_enumset1(esk1_0,esk1_0,esk1_0,k1_relat_1(esk2_0)))),k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)))|r2_hidden(esk3_3(k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)),k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)),k1_setfam_1(k2_enumset1(esk1_0,esk1_0,esk1_0,k1_relat_1(esk2_0)))),k1_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_50, c_0_80])).
cnf(c_0_87, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_88, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(X1,k1_relat_1(k5_relat_1(X2,X3)))|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_81, c_0_28])).
cnf(c_0_89, negated_conjecture, (r2_hidden(esk3_3(k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)),k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)),k1_setfam_1(k2_enumset1(esk1_0,esk1_0,esk1_0,k1_relat_1(esk2_0)))),k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)))|r2_hidden(esk3_3(k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)),k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)),k1_setfam_1(k2_enumset1(esk1_0,esk1_0,esk1_0,k1_relat_1(esk2_0)))),esk1_0)), inference(spm,[status(thm)],[c_0_82, c_0_80])).
cnf(c_0_90, plain, (k1_setfam_1(k2_enumset1(X1,X1,X1,X2))=k1_setfam_1(k2_enumset1(X3,X3,X3,k1_relat_1(k5_relat_1(k6_relat_1(X4),X5))))|~r2_hidden(esk3_3(X3,k1_relat_1(k5_relat_1(k6_relat_1(X4),X5)),k1_setfam_1(k2_enumset1(X1,X1,X1,X2))),k1_relat_1(X5))|~r2_hidden(esk3_3(X3,k1_relat_1(k5_relat_1(k6_relat_1(X4),X5)),k1_setfam_1(k2_enumset1(X1,X1,X1,X2))),X3)|~r2_hidden(esk3_3(X3,k1_relat_1(k5_relat_1(k6_relat_1(X4),X5)),k1_setfam_1(k2_enumset1(X1,X1,X1,X2))),X2)|~r2_hidden(esk3_3(X3,k1_relat_1(k5_relat_1(k6_relat_1(X4),X5)),k1_setfam_1(k2_enumset1(X1,X1,X1,X2))),X1)|~r2_hidden(esk3_3(X3,k1_relat_1(k5_relat_1(k6_relat_1(X4),X5)),k1_setfam_1(k2_enumset1(X1,X1,X1,X2))),X4)|~v1_relat_1(X5)), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_91, negated_conjecture, (r2_hidden(esk3_3(k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)),k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)),k1_setfam_1(k2_enumset1(esk1_0,esk1_0,esk1_0,k1_relat_1(esk2_0)))),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_87])])).
cnf(c_0_92, negated_conjecture, (r2_hidden(esk3_3(k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)),k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)),k1_setfam_1(k2_enumset1(esk1_0,esk1_0,esk1_0,k1_relat_1(esk2_0)))),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_29]), c_0_49]), c_0_30]), c_0_87])])).
cnf(c_0_93, negated_conjecture, (~r2_hidden(esk3_3(k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)),k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)),k1_setfam_1(k2_enumset1(esk1_0,esk1_0,esk1_0,k1_relat_1(esk2_0)))),k1_relat_1(k5_relat_1(k6_relat_1(esk1_0),esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_74]), c_0_91]), c_0_92]), c_0_87])]), c_0_65])).
cnf(c_0_94, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_84]), c_0_91]), c_0_92]), c_0_87])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 95
# Proof object clause steps            : 66
# Proof object formula steps           : 29
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 29
# Proof object simplifying inferences  : 76
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 22
# Processed clauses                    : 6927
# ...of these trivial                  : 13
# ...subsumed                          : 3646
# ...remaining for further processing  : 3268
# Other redundant clauses eliminated   : 2346
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 610
# Backward-rewritten                   : 322
# Generated clauses                    : 389384
# ...of the previous two non-trivial   : 381480
# Contextual simplify-reflections      : 95
# Paramodulations                      : 386709
# Factorizations                       : 328
# Equation resolutions                 : 2346
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2310
#    Positive orientable unit clauses  : 31
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 2274
# Current number of unprocessed clauses: 373832
# ...number of literals in the above   : 2348864
# Current number of archived formulas  : 0
# Current number of archived clauses   : 958
# Clause-clause subsumption calls (NU) : 1250906
# Rec. Clause-clause subsumption calls : 545493
# Non-unit clause-clause subsumptions  : 3802
# Unit Clause-clause subsumption calls : 10672
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 654
# BW rewrite match successes           : 25
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 14777400

# -------------------------------------------------
# User time                : 8.787 s
# System time              : 0.130 s
# Total time               : 8.916 s
# Maximum resident set size: 286716 pages
