# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t6_yellow_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>(v1_lattice3(k2_yellow_1(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>r1_tarski(k2_xboole_0(X2,X3),k10_lattice3(k2_yellow_1(X1),X2,X3)))))), file('data/MPTP/proofs/12b798c7eddb4e12ac8e74bbfa900537.E_input', t6_yellow_1)).
fof(t3_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r3_orders_2(k2_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3))))), file('data/MPTP/proofs/12b798c7eddb4e12ac8e74bbfa900537.E_input', t3_yellow_1)).
fof(t1_yellow_1, axiom, ![X1]:(u1_struct_0(k2_yellow_1(X1))=X1&u1_orders_2(k2_yellow_1(X1))=k1_yellow_1(X1)), file('data/MPTP/proofs/12b798c7eddb4e12ac8e74bbfa900537.E_input', t1_yellow_1)).
fof(t13_lattice3, axiom, ![X1]:((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k10_lattice3(X1,X2,X3)=k10_lattice3(X1,X3,X2)))), file('data/MPTP/proofs/12b798c7eddb4e12ac8e74bbfa900537.E_input', t13_lattice3)).
fof(redefinition_k13_lattice3, axiom, ![X1, X2, X3]:(((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k13_lattice3(X1,X2,X3)=k10_lattice3(X1,X2,X3)), file('data/MPTP/proofs/12b798c7eddb4e12ac8e74bbfa900537.E_input', redefinition_k13_lattice3)).
fof(t8_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X3,X2))=>r1_tarski(k2_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/12b798c7eddb4e12ac8e74bbfa900537.E_input', t8_xboole_1)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/12b798c7eddb4e12ac8e74bbfa900537.E_input', redefinition_r3_orders_2)).
fof(fc5_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&(v3_orders_2(k2_yellow_1(X1))&(v4_orders_2(k2_yellow_1(X1))&v5_orders_2(k2_yellow_1(X1))))), file('data/MPTP/proofs/12b798c7eddb4e12ac8e74bbfa900537.E_input', fc5_yellow_1)).
fof(dt_k2_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&l1_orders_2(k2_yellow_1(X1))), file('data/MPTP/proofs/12b798c7eddb4e12ac8e74bbfa900537.E_input', dt_k2_yellow_1)).
fof(t22_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k13_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))))))), file('data/MPTP/proofs/12b798c7eddb4e12ac8e74bbfa900537.E_input', t22_yellow_0)).
fof(dt_k13_lattice3, axiom, ![X1, X2, X3]:(((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k13_lattice3(X1,X2,X3),u1_struct_0(X1))), file('data/MPTP/proofs/12b798c7eddb4e12ac8e74bbfa900537.E_input', dt_k13_lattice3)).
fof(cc1_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v1_lattice3(X1)=>~(v2_struct_0(X1)))), file('data/MPTP/proofs/12b798c7eddb4e12ac8e74bbfa900537.E_input', cc1_lattice3)).
fof(c_0_12, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>(v1_lattice3(k2_yellow_1(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>r1_tarski(k2_xboole_0(X2,X3),k10_lattice3(k2_yellow_1(X1),X2,X3))))))), inference(assume_negation,[status(cth)],[t6_yellow_1])).
fof(c_0_13, plain, ![X34, X35, X36]:((~r3_orders_2(k2_yellow_1(X34),X35,X36)|r1_tarski(X35,X36)|~m1_subset_1(X36,u1_struct_0(k2_yellow_1(X34)))|~m1_subset_1(X35,u1_struct_0(k2_yellow_1(X34)))|v1_xboole_0(X34))&(~r1_tarski(X35,X36)|r3_orders_2(k2_yellow_1(X34),X35,X36)|~m1_subset_1(X36,u1_struct_0(k2_yellow_1(X34)))|~m1_subset_1(X35,u1_struct_0(k2_yellow_1(X34)))|v1_xboole_0(X34))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_yellow_1])])])])])).
fof(c_0_14, plain, ![X21]:(u1_struct_0(k2_yellow_1(X21))=X21&u1_orders_2(k2_yellow_1(X21))=k1_yellow_1(X21)), inference(variable_rename,[status(thm)],[t1_yellow_1])).
fof(c_0_15, plain, ![X15, X16, X17]:(~v5_orders_2(X15)|(~v1_lattice3(X15)|~l1_orders_2(X15))|(~m1_subset_1(X16,u1_struct_0(X15))|(~m1_subset_1(X17,u1_struct_0(X15))|k10_lattice3(X15,X16,X17)=k10_lattice3(X15,X17,X16)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_lattice3])])])).
fof(c_0_16, plain, ![X22, X23, X24]:(~v5_orders_2(X22)|(~v1_lattice3(X22)|~l1_orders_2(X22))|(~m1_subset_1(X23,u1_struct_0(X22))|~m1_subset_1(X24,u1_struct_0(X22)))|k13_lattice3(X22,X23,X24)=k10_lattice3(X22,X23,X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k13_lattice3])])).
fof(c_0_17, negated_conjecture, (~v1_xboole_0(esk1_0)&(v1_lattice3(k2_yellow_1(esk1_0))&(m1_subset_1(esk2_0,u1_struct_0(k2_yellow_1(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(k2_yellow_1(esk1_0)))&~r1_tarski(k2_xboole_0(esk2_0,esk3_0),k10_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])).
fof(c_0_18, plain, ![X12, X13, X14]:(~r1_tarski(X12,X13)|~r1_tarski(X14,X13)|r1_tarski(k2_xboole_0(X12,X14),X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_xboole_1])])).
cnf(c_0_19, plain, (r1_tarski(X2,X3)|v1_xboole_0(X1)|~r3_orders_2(k2_yellow_1(X1),X2,X3)|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, plain, (u1_struct_0(k2_yellow_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_21, plain, ![X18, X19, X20]:((~r3_orders_2(X18,X19,X20)|r1_orders_2(X18,X19,X20)|(v2_struct_0(X18)|(~v3_orders_2(X18)|~l1_orders_2(X18))|(~m1_subset_1(X19,u1_struct_0(X18))|~m1_subset_1(X20,u1_struct_0(X18)))))&(~r1_orders_2(X18,X19,X20)|r3_orders_2(X18,X19,X20)|(v2_struct_0(X18)|(~v3_orders_2(X18)|~l1_orders_2(X18))|(~m1_subset_1(X19,u1_struct_0(X18))|~m1_subset_1(X20,u1_struct_0(X18)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
fof(c_0_22, plain, ![X32]:(v1_orders_2(k2_yellow_1(X32))&(v3_orders_2(k2_yellow_1(X32))&(v4_orders_2(k2_yellow_1(X32))&v5_orders_2(k2_yellow_1(X32))))), inference(variable_rename,[status(thm)],[fc5_yellow_1])).
fof(c_0_23, plain, ![X33]:(v1_orders_2(k2_yellow_1(X33))&l1_orders_2(k2_yellow_1(X33))), inference(variable_rename,[status(thm)],[dt_k2_yellow_1])).
fof(c_0_24, plain, ![X25, X26, X27, X28, X29]:(((r1_orders_2(X25,X26,X28)|X28!=k13_lattice3(X25,X26,X27)|~m1_subset_1(X28,u1_struct_0(X25))|~m1_subset_1(X27,u1_struct_0(X25))|~m1_subset_1(X26,u1_struct_0(X25))|(~v5_orders_2(X25)|(~v1_lattice3(X25)|~l1_orders_2(X25))))&((r1_orders_2(X25,X27,X28)|X28!=k13_lattice3(X25,X26,X27)|~m1_subset_1(X28,u1_struct_0(X25))|~m1_subset_1(X27,u1_struct_0(X25))|~m1_subset_1(X26,u1_struct_0(X25))|(~v5_orders_2(X25)|(~v1_lattice3(X25)|~l1_orders_2(X25))))&(~m1_subset_1(X29,u1_struct_0(X25))|(~r1_orders_2(X25,X26,X29)|~r1_orders_2(X25,X27,X29)|r1_orders_2(X25,X28,X29))|X28!=k13_lattice3(X25,X26,X27)|~m1_subset_1(X28,u1_struct_0(X25))|~m1_subset_1(X27,u1_struct_0(X25))|~m1_subset_1(X26,u1_struct_0(X25))|(~v5_orders_2(X25)|(~v1_lattice3(X25)|~l1_orders_2(X25))))))&((m1_subset_1(esk4_4(X25,X26,X27,X28),u1_struct_0(X25))|~r1_orders_2(X25,X27,X28)|~r1_orders_2(X25,X26,X28)|X28=k13_lattice3(X25,X26,X27)|~m1_subset_1(X28,u1_struct_0(X25))|~m1_subset_1(X27,u1_struct_0(X25))|~m1_subset_1(X26,u1_struct_0(X25))|(~v5_orders_2(X25)|(~v1_lattice3(X25)|~l1_orders_2(X25))))&(((r1_orders_2(X25,X26,esk4_4(X25,X26,X27,X28))|~r1_orders_2(X25,X27,X28)|~r1_orders_2(X25,X26,X28)|X28=k13_lattice3(X25,X26,X27)|~m1_subset_1(X28,u1_struct_0(X25))|~m1_subset_1(X27,u1_struct_0(X25))|~m1_subset_1(X26,u1_struct_0(X25))|(~v5_orders_2(X25)|(~v1_lattice3(X25)|~l1_orders_2(X25))))&(r1_orders_2(X25,X27,esk4_4(X25,X26,X27,X28))|~r1_orders_2(X25,X27,X28)|~r1_orders_2(X25,X26,X28)|X28=k13_lattice3(X25,X26,X27)|~m1_subset_1(X28,u1_struct_0(X25))|~m1_subset_1(X27,u1_struct_0(X25))|~m1_subset_1(X26,u1_struct_0(X25))|(~v5_orders_2(X25)|(~v1_lattice3(X25)|~l1_orders_2(X25)))))&(~r1_orders_2(X25,X28,esk4_4(X25,X26,X27,X28))|~r1_orders_2(X25,X27,X28)|~r1_orders_2(X25,X26,X28)|X28=k13_lattice3(X25,X26,X27)|~m1_subset_1(X28,u1_struct_0(X25))|~m1_subset_1(X27,u1_struct_0(X25))|~m1_subset_1(X26,u1_struct_0(X25))|(~v5_orders_2(X25)|(~v1_lattice3(X25)|~l1_orders_2(X25))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t22_yellow_0])])])])])).
fof(c_0_25, plain, ![X9, X10, X11]:(~v5_orders_2(X9)|(~v1_lattice3(X9)|~l1_orders_2(X9))|(~m1_subset_1(X10,u1_struct_0(X9))|~m1_subset_1(X11,u1_struct_0(X9)))|m1_subset_1(k13_lattice3(X9,X10,X11),u1_struct_0(X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k13_lattice3])])).
cnf(c_0_26, plain, (k10_lattice3(X1,X2,X3)=k10_lattice3(X1,X3,X2)|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_27, plain, (k13_lattice3(X1,X2,X3)=k10_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28, negated_conjecture, (~r1_tarski(k2_xboole_0(esk2_0,esk3_0),k10_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_29, plain, (r1_tarski(k2_xboole_0(X1,X3),X2)|~r1_tarski(X1,X2)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_30, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k2_yellow_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(k2_yellow_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_32, plain, (r1_tarski(X1,X2)|v1_xboole_0(X3)|~r3_orders_2(k2_yellow_1(X3),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_20]), c_0_20])).
cnf(c_0_33, plain, (r3_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_34, plain, (v3_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_35, plain, (l1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_36, plain, (r1_orders_2(X1,X2,X3)|X3!=k13_lattice3(X1,X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_37, plain, (m1_subset_1(k13_lattice3(X1,X2,X3),u1_struct_0(X1))|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_38, plain, (v5_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_39, plain, (k10_lattice3(X1,X2,X3)=k13_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_lattice3(X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_40, negated_conjecture, (~r1_tarski(esk3_0,k10_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0))|~r1_tarski(esk2_0,k10_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_41, negated_conjecture, (m1_subset_1(esk3_0,esk1_0)), inference(rw,[status(thm)],[c_0_30, c_0_20])).
cnf(c_0_42, negated_conjecture, (m1_subset_1(esk2_0,esk1_0)), inference(rw,[status(thm)],[c_0_31, c_0_20])).
cnf(c_0_43, negated_conjecture, (v1_lattice3(k2_yellow_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_44, plain, (v2_struct_0(k2_yellow_1(X1))|r1_tarski(X2,X3)|v1_xboole_0(X1)|~r1_orders_2(k2_yellow_1(X1),X2,X3)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34]), c_0_35]), c_0_20]), c_0_20])])).
cnf(c_0_45, plain, (r1_orders_2(X1,X2,k13_lattice3(X1,X3,X2))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_lattice3(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_36]), c_0_37])).
cnf(c_0_46, plain, (m1_subset_1(k13_lattice3(k2_yellow_1(X1),X2,X3),X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v1_lattice3(k2_yellow_1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_20]), c_0_35]), c_0_38])])).
cnf(c_0_47, plain, (k13_lattice3(X1,X2,X3)=k13_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_lattice3(X1)), inference(spm,[status(thm)],[c_0_27, c_0_39])).
cnf(c_0_48, negated_conjecture, (~r1_tarski(esk3_0,k13_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0))|~r1_tarski(esk2_0,k13_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_27]), c_0_35]), c_0_38]), c_0_20]), c_0_41]), c_0_20]), c_0_42]), c_0_43])])).
cnf(c_0_49, plain, (v2_struct_0(k2_yellow_1(X1))|r1_tarski(X2,k13_lattice3(k2_yellow_1(X1),X3,X2))|v1_xboole_0(X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)|~v1_lattice3(k2_yellow_1(X1))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_35]), c_0_38]), c_0_20]), c_0_20])]), c_0_46])).
cnf(c_0_50, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_51, plain, (r1_orders_2(X1,X2,k13_lattice3(X1,X2,X3))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_lattice3(X1)), inference(spm,[status(thm)],[c_0_45, c_0_47])).
fof(c_0_52, plain, ![X31]:(~l1_orders_2(X31)|(~v1_lattice3(X31)|~v2_struct_0(X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattice3])])])).
cnf(c_0_53, negated_conjecture, (v2_struct_0(k2_yellow_1(esk1_0))|~r1_tarski(esk2_0,k13_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_41]), c_0_42]), c_0_43])]), c_0_50])).
cnf(c_0_54, plain, (v2_struct_0(k2_yellow_1(X1))|r1_tarski(X2,k13_lattice3(k2_yellow_1(X1),X2,X3))|v1_xboole_0(X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)|~v1_lattice3(k2_yellow_1(X1))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_51]), c_0_35]), c_0_38]), c_0_20]), c_0_20])]), c_0_46])).
cnf(c_0_55, plain, (~l1_orders_2(X1)|~v1_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_56, negated_conjecture, (v2_struct_0(k2_yellow_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_42]), c_0_41]), c_0_43])]), c_0_50])).
cnf(c_0_57, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_35]), c_0_43])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 58
# Proof object clause steps            : 33
# Proof object formula steps           : 25
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 47
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 28
# Processed clauses                    : 59
# ...of these trivial                  : 1
# ...subsumed                          : 17
# ...remaining for further processing  : 41
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 49
# ...of the previous two non-trivial   : 46
# Contextual simplify-reflections      : 3
# Paramodulations                      : 47
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 37
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 23
# Current number of unprocessed clauses: 15
# ...number of literals in the above   : 109
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3
# Clause-clause subsumption calls (NU) : 216
# Rec. Clause-clause subsumption calls : 31
# Non-unit clause-clause subsumptions  : 18
# Unit Clause-clause subsumption calls : 6
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3830

# -------------------------------------------------
# User time                : 0.005 s
# System time              : 0.006 s
# Total time               : 0.011 s
# Maximum resident set size: 28720 pages
