# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l1_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/12c0ea5838d94affacd3a1291175e1a4.E_input', l1_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/12c0ea5838d94affacd3a1291175e1a4.E_input', t69_enumset1)).
fof(t15_funct_1, axiom, ![X1]:(~(X1=k1_xboole_0)=>![X2]:?[X3]:((v1_relat_1(X3)&v1_funct_1(X3))&(k1_relat_1(X3)=X1&k2_relat_1(X3)=k1_tarski(X2)))), file('data/MPTP/proofs/12c0ea5838d94affacd3a1291175e1a4.E_input', t15_funct_1)).
fof(t18_funct_1, conjecture, ![X1, X2]:~((~((X1=k1_xboole_0&~(X2=k1_xboole_0)))&![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>~((X2=k1_relat_1(X3)&r1_tarski(k2_relat_1(X3),X1)))))), file('data/MPTP/proofs/12c0ea5838d94affacd3a1291175e1a4.E_input', t18_funct_1)).
fof(cc1_relat_1, axiom, ![X1]:(v1_xboole_0(X1)=>v1_relat_1(X1)), file('data/MPTP/proofs/12c0ea5838d94affacd3a1291175e1a4.E_input', cc1_relat_1)).
fof(t64_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>((k1_relat_1(X1)=k1_xboole_0|k2_relat_1(X1)=k1_xboole_0)=>X1=k1_xboole_0)), file('data/MPTP/proofs/12c0ea5838d94affacd3a1291175e1a4.E_input', t64_relat_1)).
fof(s3_funct_1__e2_25__funct_1, axiom, ![X1]:?[X2]:((v1_relat_1(X2)&v1_funct_1(X2))&(k1_relat_1(X2)=X1&![X3]:(r2_hidden(X3,X1)=>k1_funct_1(X2,X3)=k1_xboole_0))), file('data/MPTP/proofs/12c0ea5838d94affacd3a1291175e1a4.E_input', s3_funct_1__e2_25__funct_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/12c0ea5838d94affacd3a1291175e1a4.E_input', t2_xboole_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/12c0ea5838d94affacd3a1291175e1a4.E_input', fc1_xboole_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/12c0ea5838d94affacd3a1291175e1a4.E_input', d1_xboole_0)).
fof(t12_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r2_hidden(X1,k1_relat_1(X2))=>r2_hidden(k1_funct_1(X2,X1),k2_relat_1(X2)))), file('data/MPTP/proofs/12c0ea5838d94affacd3a1291175e1a4.E_input', t12_funct_1)).
fof(t60_relat_1, axiom, (k1_relat_1(k1_xboole_0)=k1_xboole_0&k2_relat_1(k1_xboole_0)=k1_xboole_0), file('data/MPTP/proofs/12c0ea5838d94affacd3a1291175e1a4.E_input', t60_relat_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/12c0ea5838d94affacd3a1291175e1a4.E_input', d5_relat_1)).
fof(c_0_13, plain, ![X11, X12]:((~r1_tarski(k1_tarski(X11),X12)|r2_hidden(X11,X12))&(~r2_hidden(X11,X12)|r1_tarski(k1_tarski(X11),X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])).
fof(c_0_14, plain, ![X33]:k2_tarski(X33,X33)=k1_tarski(X33), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_15, plain, ![X8, X9]:(((v1_relat_1(esk3_2(X8,X9))|X8=k1_xboole_0)&(v1_funct_1(esk3_2(X8,X9))|X8=k1_xboole_0))&((k1_relat_1(esk3_2(X8,X9))=X8|X8=k1_xboole_0)&(k2_relat_1(esk3_2(X8,X9))=k1_tarski(X9)|X8=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t15_funct_1])])])])])])).
fof(c_0_16, negated_conjecture, ~(![X1, X2]:~((~((X1=k1_xboole_0&~(X2=k1_xboole_0)))&![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>~((X2=k1_relat_1(X3)&r1_tarski(k2_relat_1(X3),X1))))))), inference(assume_negation,[status(cth)],[t18_funct_1])).
fof(c_0_17, plain, ![X27]:(~v1_xboole_0(X27)|v1_relat_1(X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relat_1])])).
fof(c_0_18, plain, ![X28]:((k1_relat_1(X28)!=k1_xboole_0|X28=k1_xboole_0|~v1_relat_1(X28))&(k2_relat_1(X28)!=k1_xboole_0|X28=k1_xboole_0|~v1_relat_1(X28))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t64_relat_1])])])).
fof(c_0_19, plain, ![X34, X36]:((v1_relat_1(esk8_1(X34))&v1_funct_1(esk8_1(X34)))&(k1_relat_1(esk8_1(X34))=X34&(~r2_hidden(X36,X34)|k1_funct_1(esk8_1(X34),X36)=k1_xboole_0))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[s3_funct_1__e2_25__funct_1])])])])).
cnf(c_0_20, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (k2_relat_1(esk3_2(X1,X2))=k1_tarski(X2)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_23, negated_conjecture, ![X7]:((esk1_0!=k1_xboole_0|esk2_0=k1_xboole_0)&(~v1_relat_1(X7)|~v1_funct_1(X7)|(esk2_0!=k1_relat_1(X7)|~r1_tarski(k2_relat_1(X7),esk1_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_16])])])])])).
fof(c_0_24, plain, ![X24]:r1_tarski(k1_xboole_0,X24), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_25, plain, (v1_relat_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_27, plain, (X1=k1_xboole_0|k1_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, plain, (k1_relat_1(esk8_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (v1_relat_1(esk8_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_30, plain, ![X29, X30, X31]:((~v1_xboole_0(X29)|~r2_hidden(X30,X29))&(r2_hidden(esk7_1(X31),X31)|v1_xboole_0(X31))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_31, plain, ![X25, X26]:(~v1_relat_1(X26)|~v1_funct_1(X26)|(~r2_hidden(X25,k1_relat_1(X26))|r2_hidden(k1_funct_1(X26,X25),k2_relat_1(X26)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_funct_1])])).
cnf(c_0_32, plain, (r1_tarski(k2_tarski(X1,X1),X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_33, plain, (X1=k1_xboole_0|k2_relat_1(esk3_2(X1,X2))=k2_tarski(X2,X2)), inference(rw,[status(thm)],[c_0_22, c_0_21])).
cnf(c_0_34, negated_conjecture, (~v1_relat_1(X1)|~v1_funct_1(X1)|esk2_0!=k1_relat_1(X1)|~r1_tarski(k2_relat_1(X1),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_35, plain, (k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
cnf(c_0_36, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
cnf(c_0_37, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_38, plain, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_39, plain, (v1_funct_1(esk8_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_40, plain, (esk8_1(k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])])])).
cnf(c_0_41, plain, (k1_funct_1(esk8_1(X2),X1)=k1_xboole_0|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_42, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_43, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_44, plain, (X1=k1_xboole_0|r1_tarski(k2_relat_1(esk3_2(X1,X2)),X3)|~r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_45, plain, (v1_relat_1(esk3_2(X1,X2))|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_46, plain, (v1_funct_1(esk3_2(X1,X2))|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_47, negated_conjecture, (esk2_0!=k1_xboole_0|~v1_funct_1(k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_37]), c_0_38])])).
cnf(c_0_48, plain, (v1_funct_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_49, plain, (k1_funct_1(k1_xboole_0,X1)=k1_xboole_0|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_41, c_0_40])).
fof(c_0_50, plain, ![X13, X14, X15, X17, X18, X19, X20, X22]:(((~r2_hidden(X15,X14)|r2_hidden(k4_tarski(esk4_3(X13,X14,X15),X15),X13)|X14!=k2_relat_1(X13))&(~r2_hidden(k4_tarski(X18,X17),X13)|r2_hidden(X17,X14)|X14!=k2_relat_1(X13)))&((~r2_hidden(esk5_2(X19,X20),X20)|~r2_hidden(k4_tarski(X22,esk5_2(X19,X20)),X19)|X20=k2_relat_1(X19))&(r2_hidden(esk5_2(X19,X20),X20)|r2_hidden(k4_tarski(esk6_2(X19,X20),esk5_2(X19,X20)),X19)|X20=k2_relat_1(X19)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_51, plain, (~r2_hidden(X1,k1_relat_1(X2))|~v1_xboole_0(k2_relat_1(X2))|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_52, negated_conjecture, (X1=k1_xboole_0|k1_relat_1(esk3_2(X1,X2))!=esk2_0|~r2_hidden(X2,esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_44]), c_0_45]), c_0_46])).
cnf(c_0_53, plain, (k1_relat_1(esk3_2(X1,X2))=X1|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_54, negated_conjecture, (esk2_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_48])])).
cnf(c_0_55, plain, (r2_hidden(k1_xboole_0,k1_xboole_0)|~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_49]), c_0_35]), c_0_36]), c_0_48]), c_0_38])])).
cnf(c_0_56, plain, (r2_hidden(esk5_2(X1,X2),X2)|r2_hidden(k4_tarski(esk6_2(X1,X2),esk5_2(X1,X2)),X1)|X2=k2_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_57, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_36]), c_0_35]), c_0_26]), c_0_48]), c_0_38])])).
cnf(c_0_58, negated_conjecture, (~r2_hidden(X1,esk1_0)), inference(sr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53])]), c_0_54])).
cnf(c_0_59, plain, (X1=k1_xboole_0|r2_hidden(esk5_2(k1_xboole_0,X1),X1)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_35]), c_0_57])).
cnf(c_0_60, negated_conjecture, (esk2_0=k1_xboole_0|esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_61, negated_conjecture, (esk1_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_62, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_61])]), c_0_54]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 63
# Proof object clause steps            : 38
# Proof object formula steps           : 25
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 30
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 26
# Processed clauses                    : 561
# ...of these trivial                  : 0
# ...subsumed                          : 339
# ...remaining for further processing  : 222
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 10
# Backward-rewritten                   : 8
# Generated clauses                    : 2450
# ...of the previous two non-trivial   : 1755
# Contextual simplify-reflections      : 26
# Paramodulations                      : 2444
# Factorizations                       : 2
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 176
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 163
# Current number of unprocessed clauses: 1235
# ...number of literals in the above   : 4672
# Current number of archived formulas  : 0
# Current number of archived clauses   : 45
# Clause-clause subsumption calls (NU) : 5184
# Rec. Clause-clause subsumption calls : 3755
# Non-unit clause-clause subsumptions  : 128
# Unit Clause-clause subsumption calls : 244
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 34981

# -------------------------------------------------
# User time                : 0.026 s
# System time              : 0.000 s
# Total time               : 0.026 s
# Maximum resident set size: 26304 pages
