# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t37_ordinal1, conjecture, ![X1]:~(![X2]:(r2_hidden(X2,X1)<=>v3_ordinal1(X2))), file('data/MPTP/proofs/131599cb9b68430293bed935174d89fd.E_input', t37_ordinal1)).
fof(t36_ordinal1, axiom, ![X1]:~((![X2]:(r2_hidden(X2,X1)=>v3_ordinal1(X2))&![X2]:(v3_ordinal1(X2)=>~(r1_tarski(X1,X2))))), file('data/MPTP/proofs/131599cb9b68430293bed935174d89fd.E_input', t36_ordinal1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/131599cb9b68430293bed935174d89fd.E_input', d3_tarski)).
fof(l49_zfmisc_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>r1_tarski(X1,k3_tarski(X2))), file('data/MPTP/proofs/131599cb9b68430293bed935174d89fd.E_input', l49_zfmisc_1)).
fof(antisymmetry_r2_hidden, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>~(r2_hidden(X2,X1))), file('data/MPTP/proofs/131599cb9b68430293bed935174d89fd.E_input', antisymmetry_r2_hidden)).
fof(t99_zfmisc_1, axiom, ![X1]:k3_tarski(k1_zfmisc_1(X1))=X1, file('data/MPTP/proofs/131599cb9b68430293bed935174d89fd.E_input', t99_zfmisc_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/131599cb9b68430293bed935174d89fd.E_input', t2_subset)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/131599cb9b68430293bed935174d89fd.E_input', t3_subset)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/131599cb9b68430293bed935174d89fd.E_input', fc1_subset_1)).
fof(c_0_9, negated_conjecture, ~(![X1]:~(![X2]:(r2_hidden(X2,X1)<=>v3_ordinal1(X2)))), inference(assume_negation,[status(cth)],[t37_ordinal1])).
fof(c_0_10, plain, ![X6]:(((v3_ordinal1(esk3_1(X6))|r2_hidden(esk2_1(X6),X6))&(r1_tarski(X6,esk3_1(X6))|r2_hidden(esk2_1(X6),X6)))&((v3_ordinal1(esk3_1(X6))|~v3_ordinal1(esk2_1(X6)))&(r1_tarski(X6,esk3_1(X6))|~v3_ordinal1(esk2_1(X6))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t36_ordinal1])])])])])).
fof(c_0_11, negated_conjecture, ![X5]:((~r2_hidden(X5,esk1_0)|v3_ordinal1(X5))&(~v3_ordinal1(X5)|r2_hidden(X5,esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
cnf(c_0_12, plain, (v3_ordinal1(esk3_1(X1))|~v3_ordinal1(esk2_1(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_13, negated_conjecture, (v3_ordinal1(X1)|~r2_hidden(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_14, plain, ![X9, X10, X11, X12, X13]:((~r1_tarski(X9,X10)|(~r2_hidden(X11,X9)|r2_hidden(X11,X10)))&((r2_hidden(esk4_2(X12,X13),X12)|r1_tarski(X12,X13))&(~r2_hidden(esk4_2(X12,X13),X13)|r1_tarski(X12,X13)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_15, plain, ![X15, X16]:(~r2_hidden(X15,X16)|r1_tarski(X15,k3_tarski(X16))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l49_zfmisc_1])])).
cnf(c_0_16, negated_conjecture, (r2_hidden(X1,esk1_0)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, negated_conjecture, (v3_ordinal1(esk3_1(X1))|~r2_hidden(esk2_1(X1),esk1_0)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_18, plain, (v3_ordinal1(esk3_1(X1))|r2_hidden(esk2_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, negated_conjecture, (r2_hidden(esk3_1(X1),esk1_0)|~r2_hidden(esk2_1(X1),esk1_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_22, negated_conjecture, (r2_hidden(esk3_1(X1),esk1_0)|r2_hidden(esk2_1(X1),X1)), inference(spm,[status(thm)],[c_0_16, c_0_18])).
fof(c_0_23, plain, ![X17, X18]:(~r2_hidden(X17,X18)|~r2_hidden(X18,X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[antisymmetry_r2_hidden])])])).
cnf(c_0_24, plain, (r2_hidden(X1,k3_tarski(X2))|~r2_hidden(X1,X3)|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_25, negated_conjecture, (r2_hidden(esk3_1(esk1_0),esk1_0)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_26, plain, (~r2_hidden(X1,X2)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27, negated_conjecture, (r2_hidden(esk3_1(esk1_0),k3_tarski(X1))|~r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
fof(c_0_28, plain, ![X19]:k3_tarski(k1_zfmisc_1(X19))=X19, inference(variable_rename,[status(thm)],[t99_zfmisc_1])).
cnf(c_0_29, negated_conjecture, (~r2_hidden(k3_tarski(X1),esk3_1(esk1_0))|~r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_30, plain, (k3_tarski(k1_zfmisc_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_31, plain, ![X20, X21]:(~m1_subset_1(X20,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_32, plain, ![X22, X23]:((~m1_subset_1(X22,k1_zfmisc_1(X23))|r1_tarski(X22,X23))&(~r1_tarski(X22,X23)|m1_subset_1(X22,k1_zfmisc_1(X23)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_33, plain, ![X24]:~v1_xboole_0(k1_zfmisc_1(X24)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
cnf(c_0_34, negated_conjecture, (~r2_hidden(X1,esk3_1(esk1_0))|~r2_hidden(esk1_0,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_35, negated_conjecture, (r2_hidden(esk3_1(esk1_0),X1)|~r2_hidden(esk1_0,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_27, c_0_30])).
cnf(c_0_36, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_37, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_39, negated_conjecture, (~r2_hidden(esk1_0,k1_zfmisc_1(esk3_1(esk1_0)))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40, plain, (r2_hidden(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_41, plain, (r1_tarski(X1,esk3_1(X1))|~v3_ordinal1(esk2_1(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_42, negated_conjecture, (~r1_tarski(esk1_0,esk3_1(esk1_0))), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_43, negated_conjecture, (r1_tarski(X1,esk3_1(X1))|~r2_hidden(esk2_1(X1),esk1_0)), inference(spm,[status(thm)],[c_0_41, c_0_13])).
cnf(c_0_44, plain, (r1_tarski(X1,esk3_1(X1))|r2_hidden(esk2_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_45, negated_conjecture, (~r2_hidden(esk2_1(esk1_0),esk1_0)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_46, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_44]), c_0_45]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 47
# Proof object clause steps            : 28
# Proof object formula steps           : 19
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 2
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 16
# Processed clauses                    : 78
# ...of these trivial                  : 0
# ...subsumed                          : 10
# ...remaining for further processing  : 67
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 77
# ...of the previous two non-trivial   : 72
# Contextual simplify-reflections      : 0
# Paramodulations                      : 77
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 51
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 8
#    Non-unit-clauses                  : 40
# Current number of unprocessed clauses: 26
# ...number of literals in the above   : 64
# Current number of archived formulas  : 0
# Current number of archived clauses   : 16
# Clause-clause subsumption calls (NU) : 124
# Rec. Clause-clause subsumption calls : 112
# Non-unit clause-clause subsumptions  : 7
# Unit Clause-clause subsumption calls : 63
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 10
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1944

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.003 s
# Total time               : 0.010 s
# Maximum resident set size: 28740 pages
