# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/136695be87a64f458523aadebb201f52.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/136695be87a64f458523aadebb201f52.E_input', t70_enumset1)).
fof(t95_xboole_1, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2)), file('data/MPTP/proofs/136695be87a64f458523aadebb201f52.E_input', t95_xboole_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/136695be87a64f458523aadebb201f52.E_input', t71_enumset1)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/136695be87a64f458523aadebb201f52.E_input', t2_boole)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/136695be87a64f458523aadebb201f52.E_input', t1_boole)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/136695be87a64f458523aadebb201f52.E_input', t98_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/136695be87a64f458523aadebb201f52.E_input', t100_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/136695be87a64f458523aadebb201f52.E_input', commutativity_k2_xboole_0)).
fof(t91_xboole_1, axiom, ![X1, X2, X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3)), file('data/MPTP/proofs/136695be87a64f458523aadebb201f52.E_input', t91_xboole_1)).
fof(t99_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k5_xboole_0(X1,X2),X3)=k2_xboole_0(k4_xboole_0(X1,k2_xboole_0(X2,X3)),k4_xboole_0(X2,k2_xboole_0(X1,X3))), file('data/MPTP/proofs/136695be87a64f458523aadebb201f52.E_input', t99_xboole_1)).
fof(t21_xboole_1, axiom, ![X1, X2]:k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/136695be87a64f458523aadebb201f52.E_input', t21_xboole_1)).
fof(t145_zfmisc_1, conjecture, ![X1, X2, X3]:~((~(r2_hidden(X1,X3))&(~(r2_hidden(X2,X3))&~(X3=k4_xboole_0(k2_xboole_0(X3,k2_tarski(X1,X2)),k2_tarski(X1,X2)))))), file('data/MPTP/proofs/136695be87a64f458523aadebb201f52.E_input', t145_zfmisc_1)).
fof(t144_zfmisc_1, axiom, ![X1, X2, X3]:~((~(r2_hidden(X1,X3))&(~(r2_hidden(X2,X3))&~(X3=k4_xboole_0(X3,k2_tarski(X1,X2)))))), file('data/MPTP/proofs/136695be87a64f458523aadebb201f52.E_input', t144_zfmisc_1)).
fof(c_0_14, plain, ![X4, X5]:k3_tarski(k2_tarski(X4,X5))=k2_xboole_0(X4,X5), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_15, plain, ![X6, X7]:k1_enumset1(X6,X6,X7)=k2_tarski(X6,X7), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_16, plain, ![X8, X9]:k3_xboole_0(X8,X9)=k5_xboole_0(k5_xboole_0(X8,X9),k2_xboole_0(X8,X9)), inference(variable_rename,[status(thm)],[t95_xboole_1])).
cnf(c_0_17, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_19, plain, ![X10, X11, X12]:k2_enumset1(X10,X10,X11,X12)=k1_enumset1(X10,X11,X12), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_20, plain, ![X13]:k3_xboole_0(X13,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
cnf(c_0_21, plain, (k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_23, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_24, plain, ![X14]:k2_xboole_0(X14,k1_xboole_0)=X14, inference(variable_rename,[status(thm)],[t1_boole])).
fof(c_0_25, plain, ![X15, X16]:k2_xboole_0(X15,X16)=k5_xboole_0(X15,k4_xboole_0(X16,X15)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_26, plain, ![X17, X18]:k4_xboole_0(X17,X18)=k5_xboole_0(X17,k3_xboole_0(X17,X18)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_27, plain, ![X19, X20]:k2_xboole_0(X19,X20)=k2_xboole_0(X20,X19), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_28, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, plain, (k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k2_enumset1(X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_22]), c_0_23])).
cnf(c_0_30, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_31, plain, ![X21, X22, X23]:k5_xboole_0(k5_xboole_0(X21,X22),X23)=k5_xboole_0(X21,k5_xboole_0(X22,X23)), inference(variable_rename,[status(thm)],[t91_xboole_1])).
cnf(c_0_32, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35, plain, (k5_xboole_0(k5_xboole_0(X1,k1_xboole_0),k3_tarski(k2_enumset1(X1,X1,X1,k1_xboole_0)))=k1_xboole_0), inference(rw,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_36, plain, (k3_tarski(k2_enumset1(X1,X1,X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_22]), c_0_23])).
cnf(c_0_37, plain, (k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, plain, (k3_tarski(k2_enumset1(X1,X1,X1,X2))=k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(k5_xboole_0(X2,X1),k3_tarski(k2_enumset1(X2,X2,X2,X1)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_22]), c_0_33]), c_0_23]), c_0_29])).
cnf(c_0_39, plain, (k3_tarski(k2_enumset1(X1,X1,X1,X2))=k3_tarski(k2_enumset1(X2,X2,X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_22]), c_0_22]), c_0_23]), c_0_23])).
cnf(c_0_40, plain, (k5_xboole_0(X1,k5_xboole_0(k1_xboole_0,X1))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_36]), c_0_37])).
cnf(c_0_41, plain, (k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X2,k5_xboole_0(X1,k3_tarski(k2_enumset1(X2,X2,X2,X1))))))=k3_tarski(k2_enumset1(X1,X1,X1,X2))), inference(rw,[status(thm)],[c_0_38, c_0_37])).
cnf(c_0_42, plain, (k3_tarski(k2_enumset1(k1_xboole_0,k1_xboole_0,k1_xboole_0,X1))=X1), inference(spm,[status(thm)],[c_0_36, c_0_39])).
cnf(c_0_43, plain, (k5_xboole_0(X1,k5_xboole_0(k1_xboole_0,k5_xboole_0(X1,X2)))=k5_xboole_0(k1_xboole_0,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_40]), c_0_37])).
cnf(c_0_44, plain, (k5_xboole_0(k1_xboole_0,k5_xboole_0(X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_36]), c_0_40]), c_0_42])).
cnf(c_0_45, plain, (k5_xboole_0(X1,X1)=k5_xboole_0(k1_xboole_0,k1_xboole_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_41]), c_0_42]), c_0_36]), c_0_40])).
cnf(c_0_46, plain, (k5_xboole_0(k1_xboole_0,k5_xboole_0(k1_xboole_0,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_44])).
fof(c_0_47, plain, ![X24, X25, X26]:k4_xboole_0(k5_xboole_0(X24,X25),X26)=k2_xboole_0(k4_xboole_0(X24,k2_xboole_0(X25,X26)),k4_xboole_0(X25,k2_xboole_0(X24,X26))), inference(variable_rename,[status(thm)],[t99_xboole_1])).
cnf(c_0_48, plain, (k5_xboole_0(X1,k5_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_45]), c_0_37]), c_0_46])).
cnf(c_0_49, plain, (k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(k1_xboole_0,k5_xboole_0(X1,X2))))=k1_xboole_0), inference(spm,[status(thm)],[c_0_40, c_0_37])).
fof(c_0_50, plain, ![X27, X28]:k3_xboole_0(X27,k2_xboole_0(X27,X28))=X27, inference(variable_rename,[status(thm)],[t21_xboole_1])).
cnf(c_0_51, plain, (k4_xboole_0(k5_xboole_0(X1,X2),X3)=k2_xboole_0(k4_xboole_0(X1,k2_xboole_0(X2,X3)),k4_xboole_0(X2,k2_xboole_0(X1,X3)))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_52, plain, (k5_xboole_0(X1,k5_xboole_0(k1_xboole_0,k5_xboole_0(X2,X1)))=k5_xboole_0(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_53, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_54, plain, (k5_xboole_0(k5_xboole_0(X1,X2),k5_xboole_0(k5_xboole_0(k5_xboole_0(X1,X2),X3),k3_tarski(k2_enumset1(k5_xboole_0(X1,X2),k5_xboole_0(X1,X2),k5_xboole_0(X1,X2),X3))))=k3_tarski(k2_enumset1(k5_xboole_0(X1,k5_xboole_0(k5_xboole_0(X1,k3_tarski(k2_enumset1(X2,X2,X2,X3))),k3_tarski(k2_enumset1(X1,X1,X1,k3_tarski(k2_enumset1(X2,X2,X2,X3)))))),k5_xboole_0(X1,k5_xboole_0(k5_xboole_0(X1,k3_tarski(k2_enumset1(X2,X2,X2,X3))),k3_tarski(k2_enumset1(X1,X1,X1,k3_tarski(k2_enumset1(X2,X2,X2,X3)))))),k5_xboole_0(X1,k5_xboole_0(k5_xboole_0(X1,k3_tarski(k2_enumset1(X2,X2,X2,X3))),k3_tarski(k2_enumset1(X1,X1,X1,k3_tarski(k2_enumset1(X2,X2,X2,X3)))))),k5_xboole_0(X2,k5_xboole_0(k5_xboole_0(X2,k3_tarski(k2_enumset1(X1,X1,X1,X3))),k3_tarski(k2_enumset1(X2,X2,X2,k3_tarski(k2_enumset1(X1,X1,X1,X3))))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_22]), c_0_22]), c_0_22]), c_0_33]), c_0_33]), c_0_33]), c_0_33]), c_0_23]), c_0_23]), c_0_23]), c_0_23]), c_0_29]), c_0_29]), c_0_29]), c_0_29]), c_0_29])).
cnf(c_0_55, plain, (k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_52]), c_0_44]), c_0_48])).
cnf(c_0_56, plain, (k5_xboole_0(k5_xboole_0(X1,k3_tarski(k2_enumset1(X1,X1,X1,X2))),k3_tarski(k2_enumset1(X1,X1,X1,k3_tarski(k2_enumset1(X1,X1,X1,X2)))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_22]), c_0_23]), c_0_29])).
cnf(c_0_57, plain, (k3_tarski(k2_enumset1(k5_xboole_0(X1,k5_xboole_0(X1,k5_xboole_0(k3_tarski(k2_enumset1(X2,X2,X2,X3)),k3_tarski(k2_enumset1(X1,X1,X1,k3_tarski(k2_enumset1(X2,X2,X2,X3))))))),k5_xboole_0(X1,k5_xboole_0(X1,k5_xboole_0(k3_tarski(k2_enumset1(X2,X2,X2,X3)),k3_tarski(k2_enumset1(X1,X1,X1,k3_tarski(k2_enumset1(X2,X2,X2,X3))))))),k5_xboole_0(X1,k5_xboole_0(X1,k5_xboole_0(k3_tarski(k2_enumset1(X2,X2,X2,X3)),k3_tarski(k2_enumset1(X1,X1,X1,k3_tarski(k2_enumset1(X2,X2,X2,X3))))))),k5_xboole_0(X2,k5_xboole_0(X2,k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X3)),k3_tarski(k2_enumset1(X2,X2,X2,k3_tarski(k2_enumset1(X1,X1,X1,X3)))))))))=k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X3,k3_tarski(k2_enumset1(k5_xboole_0(X1,X2),k5_xboole_0(X1,X2),k5_xboole_0(X1,X2),X3)))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_37]), c_0_37]), c_0_37]), c_0_37]), c_0_37]), c_0_37]), c_0_37]), c_0_37])).
cnf(c_0_58, plain, (k5_xboole_0(X1,k5_xboole_0(X2,X1))=X2), inference(spm,[status(thm)],[c_0_48, c_0_55])).
fof(c_0_59, negated_conjecture, ~(![X1, X2, X3]:~((~(r2_hidden(X1,X3))&(~(r2_hidden(X2,X3))&~(X3=k4_xboole_0(k2_xboole_0(X3,k2_tarski(X1,X2)),k2_tarski(X1,X2))))))), inference(assume_negation,[status(cth)],[t145_zfmisc_1])).
cnf(c_0_60, plain, (k5_xboole_0(X1,k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k3_tarski(k2_enumset1(X1,X1,X1,k3_tarski(k2_enumset1(X1,X1,X1,X2))))))=X1), inference(rw,[status(thm)],[c_0_56, c_0_37])).
cnf(c_0_61, plain, (k3_tarski(k2_enumset1(k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k3_tarski(k2_enumset1(X3,X3,X3,k3_tarski(k2_enumset1(X1,X1,X1,X2))))),k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k3_tarski(k2_enumset1(X3,X3,X3,k3_tarski(k2_enumset1(X1,X1,X1,X2))))),k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k3_tarski(k2_enumset1(X3,X3,X3,k3_tarski(k2_enumset1(X1,X1,X1,X2))))),k5_xboole_0(k3_tarski(k2_enumset1(X3,X3,X3,X2)),k3_tarski(k2_enumset1(X1,X1,X1,k3_tarski(k2_enumset1(X3,X3,X3,X2)))))))=k5_xboole_0(X3,k5_xboole_0(X1,k5_xboole_0(X3,k5_xboole_0(X1,k5_xboole_0(X2,k3_tarski(k2_enumset1(k5_xboole_0(X3,X1),k5_xboole_0(X3,X1),k5_xboole_0(X3,X1),X2)))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_48]), c_0_48]), c_0_48]), c_0_48])).
cnf(c_0_62, plain, (k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X1,X3)))=k5_xboole_0(X2,X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_58]), c_0_37])).
fof(c_0_63, negated_conjecture, (~r2_hidden(esk1_0,esk3_0)&(~r2_hidden(esk2_0,esk3_0)&esk3_0!=k4_xboole_0(k2_xboole_0(esk3_0,k2_tarski(esk1_0,esk2_0)),k2_tarski(esk1_0,esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_59])])])])).
cnf(c_0_64, plain, (k3_tarski(k2_enumset1(X1,X1,X1,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_36]), c_0_48])).
cnf(c_0_65, plain, (k3_tarski(k2_enumset1(k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k3_tarski(k2_enumset1(X3,X3,X3,k3_tarski(k2_enumset1(X1,X1,X1,X2))))),k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k3_tarski(k2_enumset1(X3,X3,X3,k3_tarski(k2_enumset1(X1,X1,X1,X2))))),k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k3_tarski(k2_enumset1(X3,X3,X3,k3_tarski(k2_enumset1(X1,X1,X1,X2))))),k5_xboole_0(k3_tarski(k2_enumset1(X3,X3,X3,X2)),k3_tarski(k2_enumset1(X1,X1,X1,k3_tarski(k2_enumset1(X3,X3,X3,X2)))))))=k5_xboole_0(X2,k3_tarski(k2_enumset1(k5_xboole_0(X3,X1),k5_xboole_0(X3,X1),k5_xboole_0(X3,X1),X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_62]), c_0_48])).
cnf(c_0_66, plain, (k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k3_tarski(k2_enumset1(X1,X1,X1,k3_tarski(k2_enumset1(X1,X1,X1,X2)))))=k5_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_48, c_0_60])).
cnf(c_0_67, plain, (k5_xboole_0(X1,X1)=k5_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_45, c_0_45])).
cnf(c_0_68, negated_conjecture, (esk3_0!=k4_xboole_0(k2_xboole_0(esk3_0,k2_tarski(esk1_0,esk2_0)),k2_tarski(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_69, plain, (k5_xboole_0(X1,k3_tarski(k2_enumset1(k5_xboole_0(X2,X2),k5_xboole_0(X2,X2),k5_xboole_0(X2,X2),X1)))=k5_xboole_0(X2,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_66])).
cnf(c_0_70, plain, (k5_xboole_0(X1,k5_xboole_0(X2,X2))=X1), inference(spm,[status(thm)],[c_0_48, c_0_67])).
fof(c_0_71, plain, ![X32, X33, X34]:(r2_hidden(X32,X34)|(r2_hidden(X33,X34)|X34=k4_xboole_0(X34,k2_tarski(X32,X33)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t144_zfmisc_1])])])).
cnf(c_0_72, negated_conjecture, (esk3_0!=k5_xboole_0(k3_tarski(k2_enumset1(esk3_0,esk3_0,esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0))),k5_xboole_0(k5_xboole_0(k3_tarski(k2_enumset1(esk3_0,esk3_0,esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0))),k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0)),k3_tarski(k2_enumset1(k3_tarski(k2_enumset1(esk3_0,esk3_0,esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0))),k3_tarski(k2_enumset1(esk3_0,esk3_0,esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0))),k3_tarski(k2_enumset1(esk3_0,esk3_0,esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0))),k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_18]), c_0_18]), c_0_22]), c_0_33]), c_0_23]), c_0_23]), c_0_23]), c_0_23]), c_0_23]), c_0_29])).
cnf(c_0_73, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_36]), c_0_70])).
cnf(c_0_74, plain, (r2_hidden(X1,X2)|r2_hidden(X3,X2)|X2=k4_xboole_0(X2,k2_tarski(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_75, negated_conjecture, (k5_xboole_0(k3_tarski(k2_enumset1(esk3_0,esk3_0,esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0))),k5_xboole_0(k3_tarski(k2_enumset1(esk3_0,esk3_0,esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0))),k5_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k3_tarski(k2_enumset1(k3_tarski(k2_enumset1(esk3_0,esk3_0,esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0))),k3_tarski(k2_enumset1(esk3_0,esk3_0,esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0))),k3_tarski(k2_enumset1(esk3_0,esk3_0,esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0))),k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0))))))!=esk3_0), inference(rw,[status(thm)],[c_0_72, c_0_37])).
cnf(c_0_76, plain, (k5_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),k3_tarski(k2_enumset1(X1,X1,X1,k3_tarski(k2_enumset1(X1,X1,X1,X2)))))=k1_xboole_0), inference(rw,[status(thm)],[c_0_66, c_0_73])).
cnf(c_0_77, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_70, c_0_73])).
cnf(c_0_78, plain, (X2=k5_xboole_0(X2,k5_xboole_0(k5_xboole_0(X2,k2_enumset1(X1,X1,X1,X3)),k3_tarski(k2_enumset1(X2,X2,X2,k2_enumset1(X1,X1,X1,X3)))))|r2_hidden(X3,X2)|r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_18]), c_0_33]), c_0_23]), c_0_29])).
cnf(c_0_79, negated_conjecture, (k5_xboole_0(k3_tarski(k2_enumset1(esk3_0,esk3_0,esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0))),k5_xboole_0(k3_tarski(k2_enumset1(esk3_0,esk3_0,esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0))),k5_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k3_tarski(k2_enumset1(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k3_tarski(k2_enumset1(esk3_0,esk3_0,esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0))))))))!=esk3_0), inference(rw,[status(thm)],[c_0_75, c_0_39])).
cnf(c_0_80, plain, (k3_tarski(k2_enumset1(X1,X1,X1,k3_tarski(k2_enumset1(X1,X1,X1,X2))))=k3_tarski(k2_enumset1(X1,X1,X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_76]), c_0_77])).
cnf(c_0_81, plain, (k5_xboole_0(X1,k5_xboole_0(X1,k5_xboole_0(k2_enumset1(X2,X2,X2,X3),k3_tarski(k2_enumset1(X1,X1,X1,k2_enumset1(X2,X2,X2,X3))))))=X1|r2_hidden(X2,X1)|r2_hidden(X3,X1)), inference(rw,[status(thm)],[c_0_78, c_0_37])).
cnf(c_0_82, negated_conjecture, (k5_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k3_tarski(k2_enumset1(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k3_tarski(k2_enumset1(esk3_0,esk3_0,esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0))))))!=esk3_0), inference(rw,[status(thm)],[c_0_79, c_0_48])).
cnf(c_0_83, plain, (k3_tarski(k2_enumset1(X1,X1,X1,k3_tarski(k2_enumset1(X2,X2,X2,X1))))=k3_tarski(k2_enumset1(X2,X2,X2,X1))), inference(spm,[status(thm)],[c_0_80, c_0_39])).
cnf(c_0_84, plain, (k5_xboole_0(k2_enumset1(X1,X1,X1,X2),k3_tarski(k2_enumset1(X3,X3,X3,k2_enumset1(X1,X1,X1,X2))))=X3|r2_hidden(X2,X3)|r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_81, c_0_48])).
cnf(c_0_85, negated_conjecture, (k5_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0),k3_tarski(k2_enumset1(esk3_0,esk3_0,esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0))))!=esk3_0), inference(rw,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_86, plain, (k3_tarski(k2_enumset1(X1,X1,X1,k2_enumset1(X2,X2,X2,X3)))=k5_xboole_0(k2_enumset1(X2,X2,X2,X3),X1)|r2_hidden(X2,X1)|r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_48, c_0_84])).
cnf(c_0_87, negated_conjecture, (~r2_hidden(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_88, negated_conjecture, (~r2_hidden(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_89, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_55]), c_0_58])]), c_0_87]), c_0_88]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 90
# Proof object clause steps            : 61
# Proof object formula steps           : 29
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 94
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 11
# Processed clauses                    : 474
# ...of these trivial                  : 74
# ...subsumed                          : 313
# ...remaining for further processing  : 87
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 7
# Backward-rewritten                   : 30
# Generated clauses                    : 3252
# ...of the previous two non-trivial   : 2625
# Contextual simplify-reflections      : 0
# Paramodulations                      : 3246
# Factorizations                       : 6
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 39
#    Positive orientable unit clauses  : 20
#    Positive unorientable unit clauses: 5
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 10
# Current number of unprocessed clauses: 1867
# ...number of literals in the above   : 2802
# Current number of archived formulas  : 0
# Current number of archived clauses   : 53
# Clause-clause subsumption calls (NU) : 120
# Rec. Clause-clause subsumption calls : 99
# Non-unit clause-clause subsumptions  : 55
# Unit Clause-clause subsumption calls : 75
# Rewrite failures with RHS unbound    : 28
# BW rewrite match attempts            : 573
# BW rewrite match successes           : 201
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 232662

# -------------------------------------------------
# User time                : 0.054 s
# System time              : 0.003 s
# Total time               : 0.058 s
# Maximum resident set size: 28680 pages
