# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t49_tex_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v1_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_tex_2(X2,X1)<=>~(v1_subset_1(X2,u1_struct_0(X1)))))), file('data/MPTP/proofs/1396482ba55e4f8e804a471a6a9d9959.E_input', t49_tex_2)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/1396482ba55e4f8e804a471a6a9d9959.E_input', l3_subset_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/1396482ba55e4f8e804a471a6a9d9959.E_input', d3_tarski)).
fof(d7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(v1_subset_1(X2,X1)<=>~(X2=X1))), file('data/MPTP/proofs/1396482ba55e4f8e804a471a6a9d9959.E_input', d7_subset_1)).
fof(d7_tex_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_tex_2(X2,X1)<=>(v2_tex_2(X2,X1)&![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((v2_tex_2(X3,X1)&r1_tarski(X2,X3))=>X2=X3)))))), file('data/MPTP/proofs/1396482ba55e4f8e804a471a6a9d9959.E_input', d7_tex_2)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/1396482ba55e4f8e804a471a6a9d9959.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/1396482ba55e4f8e804a471a6a9d9959.E_input', d4_subset_1)).
fof(t43_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v1_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>v2_tex_2(X2,X1))), file('data/MPTP/proofs/1396482ba55e4f8e804a471a6a9d9959.E_input', t43_tex_2)).
fof(cc1_tdlat_3, axiom, ![X1]:(l1_pre_topc(X1)=>(v1_tdlat_3(X1)=>v2_pre_topc(X1))), file('data/MPTP/proofs/1396482ba55e4f8e804a471a6a9d9959.E_input', cc1_tdlat_3)).
fof(dt_k2_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/1396482ba55e4f8e804a471a6a9d9959.E_input', dt_k2_struct_0)).
fof(fc12_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>~(v1_subset_1(k2_struct_0(X1),u1_struct_0(X1)))), file('data/MPTP/proofs/1396482ba55e4f8e804a471a6a9d9959.E_input', fc12_struct_0)).
fof(t17_tdlat_3, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>(v1_tdlat_3(X1)<=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>v3_pre_topc(X2,X1)))), file('data/MPTP/proofs/1396482ba55e4f8e804a471a6a9d9959.E_input', t17_tdlat_3)).
fof(t48_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v1_tops_1(X2,X1)&v2_tex_2(X2,X1))=>v3_tex_2(X2,X1)))), file('data/MPTP/proofs/1396482ba55e4f8e804a471a6a9d9959.E_input', t48_tex_2)).
fof(d2_tops_3, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v1_tops_1(X2,X1)<=>k2_pre_topc(X1,X2)=u1_struct_0(X1)))), file('data/MPTP/proofs/1396482ba55e4f8e804a471a6a9d9959.E_input', d2_tops_3)).
fof(fc14_subset_1, axiom, ![X1]:~(v1_subset_1(k2_subset_1(X1),X1)), file('data/MPTP/proofs/1396482ba55e4f8e804a471a6a9d9959.E_input', fc14_subset_1)).
fof(d6_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v4_pre_topc(X2,X1)<=>v3_pre_topc(k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2),X1)))), file('data/MPTP/proofs/1396482ba55e4f8e804a471a6a9d9959.E_input', d6_pre_topc)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/1396482ba55e4f8e804a471a6a9d9959.E_input', dt_l1_pre_topc)).
fof(dt_k7_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>m1_subset_1(k7_subset_1(X1,X2,X3),k1_zfmisc_1(X1))), file('data/MPTP/proofs/1396482ba55e4f8e804a471a6a9d9959.E_input', dt_k7_subset_1)).
fof(t52_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v4_pre_topc(X2,X1)=>k2_pre_topc(X1,X2)=X2)&((v2_pre_topc(X1)&k2_pre_topc(X1,X2)=X2)=>v4_pre_topc(X2,X1))))), file('data/MPTP/proofs/1396482ba55e4f8e804a471a6a9d9959.E_input', t52_pre_topc)).
fof(c_0_19, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v1_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_tex_2(X2,X1)<=>~(v1_subset_1(X2,u1_struct_0(X1))))))), inference(assume_negation,[status(cth)],[t49_tex_2])).
fof(c_0_20, plain, ![X32, X33, X34]:(~m1_subset_1(X33,k1_zfmisc_1(X32))|(~r2_hidden(X34,X33)|r2_hidden(X34,X32))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
fof(c_0_21, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&(v1_tdlat_3(esk1_0)&l1_pre_topc(esk1_0))))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&((~v3_tex_2(esk2_0,esk1_0)|v1_subset_1(esk2_0,u1_struct_0(esk1_0)))&(v3_tex_2(esk2_0,esk1_0)|~v1_subset_1(esk2_0,u1_struct_0(esk1_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_19])])])])).
fof(c_0_22, plain, ![X15, X16, X17, X18, X19]:((~r1_tarski(X15,X16)|(~r2_hidden(X17,X15)|r2_hidden(X17,X16)))&((r2_hidden(esk4_2(X18,X19),X18)|r1_tarski(X18,X19))&(~r2_hidden(esk4_2(X18,X19),X19)|r1_tarski(X18,X19)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_23, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_25, plain, ![X13, X14]:((~v1_subset_1(X14,X13)|X14!=X13|~m1_subset_1(X14,k1_zfmisc_1(X13)))&(X14=X13|v1_subset_1(X14,X13)|~m1_subset_1(X14,k1_zfmisc_1(X13)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_subset_1])])])])).
fof(c_0_26, plain, ![X24, X25, X26]:(((v2_tex_2(X25,X24)|~v3_tex_2(X25,X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_pre_topc(X24))&(~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X24)))|(~v2_tex_2(X26,X24)|~r1_tarski(X25,X26)|X25=X26)|~v3_tex_2(X25,X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_pre_topc(X24)))&((m1_subset_1(esk5_2(X24,X25),k1_zfmisc_1(u1_struct_0(X24)))|~v2_tex_2(X25,X24)|v3_tex_2(X25,X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_pre_topc(X24))&(((v2_tex_2(esk5_2(X24,X25),X24)|~v2_tex_2(X25,X24)|v3_tex_2(X25,X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_pre_topc(X24))&(r1_tarski(X25,esk5_2(X24,X25))|~v2_tex_2(X25,X24)|v3_tex_2(X25,X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_pre_topc(X24)))&(X25!=esk5_2(X24,X25)|~v2_tex_2(X25,X24)|v3_tex_2(X25,X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_pre_topc(X24))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_tex_2])])])])])).
fof(c_0_27, plain, ![X35]:m1_subset_1(k2_subset_1(X35),k1_zfmisc_1(X35)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_28, plain, ![X9]:k2_subset_1(X9)=X9, inference(variable_rename,[status(thm)],[d4_subset_1])).
cnf(c_0_29, plain, (r1_tarski(X1,X2)|~r2_hidden(esk4_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_31, plain, (X1=X2|v1_subset_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_32, plain, ![X11, X12]:(v2_struct_0(X11)|(~v2_pre_topc(X11)|(~v1_tdlat_3(X11)|~l1_pre_topc(X11)))|(~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|v2_tex_2(X12,X11))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t43_tex_2])])])])).
fof(c_0_33, plain, ![X31]:(~l1_pre_topc(X31)|(~v1_tdlat_3(X31)|v2_pre_topc(X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_tdlat_3])])).
cnf(c_0_34, plain, (X3=X1|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v2_tex_2(X1,X2)|~r1_tarski(X3,X1)|~v3_tex_2(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_36, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~r2_hidden(esk4_2(X1,u1_struct_0(esk1_0)),esk2_0)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_39, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_40, negated_conjecture, (v3_tex_2(esk2_0,esk1_0)|~v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_41, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_31, c_0_24])).
cnf(c_0_42, plain, (v2_struct_0(X1)|v2_tex_2(X2,X1)|~v2_pre_topc(X1)|~v1_tdlat_3(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43, plain, (v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_44, plain, ![X43]:(~l1_struct_0(X43)|m1_subset_1(k2_struct_0(X43),k1_zfmisc_1(u1_struct_0(X43)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_struct_0])])).
fof(c_0_45, plain, ![X10]:(~l1_struct_0(X10)|~v1_subset_1(k2_struct_0(X10),u1_struct_0(X10))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc12_struct_0])])])).
fof(c_0_46, plain, ![X6, X7]:((~v1_tdlat_3(X6)|(~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|v3_pre_topc(X7,X6))|(~v2_pre_topc(X6)|~l1_pre_topc(X6)))&((m1_subset_1(esk3_1(X6),k1_zfmisc_1(u1_struct_0(X6)))|v1_tdlat_3(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6)))&(~v3_pre_topc(esk3_1(X6),X6)|v1_tdlat_3(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t17_tdlat_3])])])])])).
cnf(c_0_47, negated_conjecture, (X1=esk2_0|~r1_tarski(esk2_0,X1)|~v2_tex_2(X1,esk1_0)|~v3_tex_2(esk2_0,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_24]), c_0_35])])).
cnf(c_0_48, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_49, negated_conjecture, (r1_tarski(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_50, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|v3_tex_2(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_51, plain, (v2_tex_2(X1,X2)|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v1_tdlat_3(X2)), inference(csr,[status(thm)],[c_0_42, c_0_43])).
fof(c_0_52, plain, ![X39, X40]:(v2_struct_0(X39)|(~v2_pre_topc(X39)|~l1_pre_topc(X39))|(~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X39)))|(~v1_tops_1(X40,X39)|~v2_tex_2(X40,X39)|v3_tex_2(X40,X39)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t48_tex_2])])])])).
cnf(c_0_53, negated_conjecture, (v1_tdlat_3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_54, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_55, plain, ![X21, X22]:((~v1_tops_1(X22,X21)|k2_pre_topc(X21,X22)=u1_struct_0(X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))|~l1_pre_topc(X21))&(k2_pre_topc(X21,X22)!=u1_struct_0(X21)|v1_tops_1(X22,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))|~l1_pre_topc(X21))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tops_3])])])])).
fof(c_0_56, plain, ![X23]:~v1_subset_1(k2_subset_1(X23),X23), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc14_subset_1])])).
fof(c_0_57, plain, ![X41, X42]:((~v4_pre_topc(X42,X41)|v3_pre_topc(k7_subset_1(u1_struct_0(X41),k2_struct_0(X41),X42),X41)|~m1_subset_1(X42,k1_zfmisc_1(u1_struct_0(X41)))|~l1_pre_topc(X41))&(~v3_pre_topc(k7_subset_1(u1_struct_0(X41),k2_struct_0(X41),X42),X41)|v4_pre_topc(X42,X41)|~m1_subset_1(X42,k1_zfmisc_1(u1_struct_0(X41)))|~l1_pre_topc(X41))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_pre_topc])])])])).
cnf(c_0_58, plain, (m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_59, plain, (~l1_struct_0(X1)|~v1_subset_1(k2_struct_0(X1),u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_60, plain, ![X38]:(~l1_pre_topc(X38)|l1_struct_0(X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_61, plain, (v3_pre_topc(X2,X1)|~v1_tdlat_3(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_62, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|~v2_tex_2(u1_struct_0(esk1_0),esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])]), c_0_50])).
cnf(c_0_63, plain, (v2_tex_2(u1_struct_0(X1),X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(spm,[status(thm)],[c_0_51, c_0_48])).
cnf(c_0_64, plain, (v2_struct_0(X1)|v3_tex_2(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_1(X2,X1)|~v2_tex_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_65, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_66, negated_conjecture, (v2_tex_2(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_24]), c_0_35]), c_0_53])]), c_0_54])).
cnf(c_0_67, plain, (v1_tops_1(X2,X1)|k2_pre_topc(X1,X2)!=u1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_68, plain, (~v1_subset_1(k2_subset_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_69, plain, (v4_pre_topc(X2,X1)|~v3_pre_topc(k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_70, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_58]), c_0_59])).
cnf(c_0_71, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_72, plain, (v3_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v1_tdlat_3(X2)), inference(csr,[status(thm)],[c_0_61, c_0_43])).
cnf(c_0_73, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_35]), c_0_53])]), c_0_54])).
fof(c_0_74, plain, ![X28, X29, X30]:(~m1_subset_1(X29,k1_zfmisc_1(X28))|m1_subset_1(k7_subset_1(X28,X29,X30),k1_zfmisc_1(X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_subset_1])])).
fof(c_0_75, plain, ![X36, X37]:((~v4_pre_topc(X37,X36)|k2_pre_topc(X36,X37)=X37|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X36)))|~l1_pre_topc(X36))&(~v2_pre_topc(X36)|k2_pre_topc(X36,X37)!=X37|v4_pre_topc(X37,X36)|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X36)))|~l1_pre_topc(X36))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t52_pre_topc])])])])).
cnf(c_0_76, negated_conjecture, (v3_tex_2(esk2_0,esk1_0)|~v1_tops_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_24]), c_0_35]), c_0_65])]), c_0_54]), c_0_66])])).
cnf(c_0_77, negated_conjecture, (v1_tops_1(esk2_0,esk1_0)|k2_pre_topc(esk1_0,esk2_0)!=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_24]), c_0_35])])).
cnf(c_0_78, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~v3_tex_2(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_79, plain, (~v1_subset_1(X1,X1)), inference(rw,[status(thm)],[c_0_68, c_0_37])).
cnf(c_0_80, plain, (v4_pre_topc(X1,X2)|~v3_pre_topc(k7_subset_1(u1_struct_0(X2),u1_struct_0(X2),X1),X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_71])).
cnf(c_0_81, negated_conjecture, (v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_35]), c_0_53])])).
cnf(c_0_82, plain, (m1_subset_1(k7_subset_1(X2,X1,X3),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_83, plain, (k2_pre_topc(X2,X1)=X1|~v4_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_84, negated_conjecture, (v3_tex_2(esk2_0,esk1_0)|k2_pre_topc(esk1_0,esk2_0)!=u1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_85, negated_conjecture, (~v3_tex_2(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_73]), c_0_79])).
cnf(c_0_86, negated_conjecture, (v4_pre_topc(X1,esk1_0)|~v3_pre_topc(k7_subset_1(esk2_0,esk2_0,X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_73]), c_0_35])])).
cnf(c_0_87, negated_conjecture, (v3_pre_topc(k7_subset_1(esk2_0,X1,X2),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_88, negated_conjecture, (k2_pre_topc(esk1_0,X1)=X1|~v4_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_73]), c_0_35])])).
cnf(c_0_89, negated_conjecture, (k2_pre_topc(esk1_0,esk2_0)!=esk2_0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_84, c_0_73]), c_0_85])).
cnf(c_0_90, negated_conjecture, (v4_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_48])])).
cnf(c_0_91, negated_conjecture, (~v4_pre_topc(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_48]), c_0_89])).
cnf(c_0_92, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_48]), c_0_91]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 93
# Proof object clause steps            : 54
# Proof object formula steps           : 39
# Proof object conjectures             : 31
# Proof object clause conjectures      : 28
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 22
# Proof object simplifying inferences  : 42
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 38
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 37
# Processed clauses                    : 132
# ...of these trivial                  : 1
# ...subsumed                          : 32
# ...remaining for further processing  : 99
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 14
# Generated clauses                    : 192
# ...of the previous two non-trivial   : 164
# Contextual simplify-reflections      : 12
# Paramodulations                      : 191
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 80
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 65
# Current number of unprocessed clauses: 66
# ...number of literals in the above   : 321
# Current number of archived formulas  : 0
# Current number of archived clauses   : 19
# Clause-clause subsumption calls (NU) : 1258
# Rec. Clause-clause subsumption calls : 439
# Non-unit clause-clause subsumptions  : 23
# Unit Clause-clause subsumption calls : 109
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5703

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28792 pages
