# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t113_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>(v1_funct_1(k7_tmap_1(X1,X2))&(v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))&(v5_pre_topc(k7_tmap_1(X1,X2),X1,k6_tmap_1(X1,X2))&m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2))))))))))), file('data/MPTP/proofs/13b977d209a54c1894fbe4f43a3be30d.E_input', t113_tmap_1)).
fof(t104_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(u1_struct_0(k6_tmap_1(X1,X2))=u1_struct_0(X1)&u1_pre_topc(k6_tmap_1(X1,X2))=k5_tmap_1(X1,X2)))), file('data/MPTP/proofs/13b977d209a54c1894fbe4f43a3be30d.E_input', t104_tmap_1)).
fof(t52_tops_2, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(l1_struct_0(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>((k2_struct_0(X2)=k1_xboole_0=>k2_struct_0(X1)=k1_xboole_0)=>k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,k2_struct_0(X2))=k2_struct_0(X1))))), file('data/MPTP/proofs/13b977d209a54c1894fbe4f43a3be30d.E_input', t52_tops_2)).
fof(dt_k6_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_pre_topc(k6_tmap_1(X1,X2))&(v2_pre_topc(k6_tmap_1(X1,X2))&l1_pre_topc(k6_tmap_1(X1,X2))))), file('data/MPTP/proofs/13b977d209a54c1894fbe4f43a3be30d.E_input', dt_k6_tmap_1)).
fof(dt_k7_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_funct_1(k7_tmap_1(X1,X2))&(v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))&m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))))), file('data/MPTP/proofs/13b977d209a54c1894fbe4f43a3be30d.E_input', dt_k7_tmap_1)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/13b977d209a54c1894fbe4f43a3be30d.E_input', dt_l1_pre_topc)).
fof(redefinition_k8_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k8_relset_1(X1,X2,X3,X4)=k10_relat_1(X3,X4)), file('data/MPTP/proofs/13b977d209a54c1894fbe4f43a3be30d.E_input', redefinition_k8_relset_1)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/13b977d209a54c1894fbe4f43a3be30d.E_input', d3_struct_0)).
fof(t171_funct_2, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k8_relset_1(X1,X1,k6_partfun1(X1),X2)=X2), file('data/MPTP/proofs/13b977d209a54c1894fbe4f43a3be30d.E_input', t171_funct_2)).
fof(d10_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k7_tmap_1(X1,X2)=k6_partfun1(u1_struct_0(X1)))), file('data/MPTP/proofs/13b977d209a54c1894fbe4f43a3be30d.E_input', d10_tmap_1)).
fof(t105_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k6_tmap_1(X1,X2))))=>(X3=X2=>v3_pre_topc(X3,k6_tmap_1(X1,X2)))))), file('data/MPTP/proofs/13b977d209a54c1894fbe4f43a3be30d.E_input', t105_tmap_1)).
fof(t55_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>((k2_struct_0(X2)=k1_xboole_0=>k2_struct_0(X1)=k1_xboole_0)=>(v5_pre_topc(X3,X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(v3_pre_topc(X4,X2)=>v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X1)))))))), file('data/MPTP/proofs/13b977d209a54c1894fbe4f43a3be30d.E_input', t55_tops_2)).
fof(t62_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((v2_pre_topc(X2)&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_struct_0(X2))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_struct_0(X2))))))=>(X3=X4=>(v5_pre_topc(X3,X1,X2)<=>v5_pre_topc(X4,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),X2))))))), file('data/MPTP/proofs/13b977d209a54c1894fbe4f43a3be30d.E_input', t62_pre_topc)).
fof(t106_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=k6_tmap_1(X1,X2)))), file('data/MPTP/proofs/13b977d209a54c1894fbe4f43a3be30d.E_input', t106_tmap_1)).
fof(c_0_14, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>(v1_funct_1(k7_tmap_1(X1,X2))&(v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))&(v5_pre_topc(k7_tmap_1(X1,X2),X1,k6_tmap_1(X1,X2))&m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))))))))), inference(assume_negation,[status(cth)],[t113_tmap_1])).
fof(c_0_15, plain, ![X7, X8]:((u1_struct_0(k6_tmap_1(X7,X8))=u1_struct_0(X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&(u1_pre_topc(k6_tmap_1(X7,X8))=k5_tmap_1(X7,X8)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t104_tmap_1])])])])])).
fof(c_0_16, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&((~v3_pre_topc(esk2_0,esk1_0)|(~v1_funct_1(k7_tmap_1(esk1_0,esk2_0))|(~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))|(~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))))))))&((v1_funct_1(k7_tmap_1(esk1_0,esk2_0))|v3_pre_topc(esk2_0,esk1_0))&((v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))|v3_pre_topc(esk2_0,esk1_0))&((v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|v3_pre_topc(esk2_0,esk1_0))&(m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))))|v3_pre_topc(esk2_0,esk1_0)))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])])).
fof(c_0_17, plain, ![X9, X10, X11]:((k2_struct_0(X10)=k1_xboole_0|k8_relset_1(u1_struct_0(X9),u1_struct_0(X10),X11,k2_struct_0(X10))=k2_struct_0(X9)|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X10))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X10))))))|~l1_struct_0(X10)|~l1_struct_0(X9))&(k2_struct_0(X9)!=k1_xboole_0|k8_relset_1(u1_struct_0(X9),u1_struct_0(X10),X11,k2_struct_0(X10))=k2_struct_0(X9)|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X10))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X10))))))|~l1_struct_0(X10)|~l1_struct_0(X9))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t52_tops_2])])])])).
cnf(c_0_18, plain, (u1_struct_0(k6_tmap_1(X1,X2))=u1_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_23, plain, ![X12, X13]:((v1_pre_topc(k6_tmap_1(X12,X13))|(v2_struct_0(X12)|(~v2_pre_topc(X12)|~l1_pre_topc(X12))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))))&((v2_pre_topc(k6_tmap_1(X12,X13))|(v2_struct_0(X12)|(~v2_pre_topc(X12)|~l1_pre_topc(X12))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))))&(l1_pre_topc(k6_tmap_1(X12,X13))|(v2_struct_0(X12)|(~v2_pre_topc(X12)|~l1_pre_topc(X12))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_tmap_1])])])])).
fof(c_0_24, plain, ![X14, X15]:((v1_funct_1(k7_tmap_1(X14,X15))|(v2_struct_0(X14)|(~v2_pre_topc(X14)|~l1_pre_topc(X14))|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X14)))))&((v1_funct_2(k7_tmap_1(X14,X15),u1_struct_0(X14),u1_struct_0(k6_tmap_1(X14,X15)))|(v2_struct_0(X14)|(~v2_pre_topc(X14)|~l1_pre_topc(X14))|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X14)))))&(m1_subset_1(k7_tmap_1(X14,X15),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X14),u1_struct_0(k6_tmap_1(X14,X15)))))|(v2_struct_0(X14)|(~v2_pre_topc(X14)|~l1_pre_topc(X14))|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X14))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_tmap_1])])])])).
cnf(c_0_25, plain, (k2_struct_0(X1)=k1_xboole_0|k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),X3,k2_struct_0(X1))=k2_struct_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~l1_struct_0(X1)|~l1_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, negated_conjecture, (u1_struct_0(k6_tmap_1(esk1_0,esk2_0))=u1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20]), c_0_21])]), c_0_22])).
fof(c_0_27, plain, ![X16]:(~l1_pre_topc(X16)|l1_struct_0(X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_28, plain, (l1_pre_topc(k6_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_29, plain, ![X17, X18, X19, X20]:(~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(X17,X18)))|k8_relset_1(X17,X18,X19,X20)=k10_relat_1(X19,X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_relset_1])])).
fof(c_0_30, plain, ![X21]:(~l1_struct_0(X21)|k2_struct_0(X21)=u1_struct_0(X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_31, plain, (v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, plain, (v1_funct_1(k7_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, plain, (m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(X1),X2,k2_struct_0(X1))=k2_struct_0(k6_tmap_1(esk1_0,esk2_0))|k2_struct_0(X1)=k1_xboole_0|~l1_struct_0(k6_tmap_1(esk1_0,esk2_0))|~l1_struct_0(X1)|~v1_funct_2(X2,u1_struct_0(esk1_0),u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_35, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36, negated_conjecture, (l1_pre_topc(k6_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_19]), c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_37, plain, (k8_relset_1(X2,X3,X1,X4)=k10_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, plain, (k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,k2_struct_0(X2))=k2_struct_0(X1)|k2_struct_0(X1)!=k1_xboole_0|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_struct_0(X2)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_39, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40, negated_conjecture, (v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_19]), c_0_26]), c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_41, negated_conjecture, (v1_funct_1(k7_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_19]), c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_42, negated_conjecture, (m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_19]), c_0_26]), c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_43, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(X1),X2,k2_struct_0(X1))=k2_struct_0(k6_tmap_1(esk1_0,esk2_0))|k2_struct_0(X1)=k1_xboole_0|~l1_struct_0(X1)|~v1_funct_2(X2,u1_struct_0(esk1_0),u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])])).
cnf(c_0_44, plain, (k2_struct_0(X1)=k10_relat_1(X2,k2_struct_0(X3))|k2_struct_0(X1)!=k1_xboole_0|~l1_struct_0(X3)|~l1_struct_0(X1)|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_45, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_39, c_0_35])).
cnf(c_0_46, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,esk2_0),k2_struct_0(esk1_0))=k2_struct_0(esk1_0)|k2_struct_0(esk1_0)=k1_xboole_0|~l1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_40]), c_0_41]), c_0_42])])).
cnf(c_0_47, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,esk2_0),k2_struct_0(esk1_0))=k2_struct_0(k6_tmap_1(esk1_0,esk2_0))|k2_struct_0(esk1_0)=k1_xboole_0|~l1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_40]), c_0_41]), c_0_42])])).
fof(c_0_48, plain, ![X22, X23]:(~m1_subset_1(X23,k1_zfmisc_1(X22))|k8_relset_1(X22,X22,k6_partfun1(X22),X23)=X23), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t171_funct_2])])).
fof(c_0_49, plain, ![X24, X25]:(v2_struct_0(X24)|(~v2_pre_topc(X24)|~l1_pre_topc(X24))|(~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|k7_tmap_1(X24,X25)=k6_partfun1(u1_struct_0(X24)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_tmap_1])])])])).
cnf(c_0_50, plain, (u1_struct_0(X1)=k10_relat_1(X2,k2_struct_0(X3))|u1_struct_0(X1)!=k1_xboole_0|~l1_struct_0(X3)|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_35])).
cnf(c_0_51, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1,k2_struct_0(k6_tmap_1(esk1_0,esk2_0)))=k2_struct_0(k6_tmap_1(esk1_0,esk2_0))|k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|~l1_struct_0(k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(spm,[status(thm)],[c_0_43, c_0_26])).
cnf(c_0_52, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,esk2_0),k2_struct_0(esk1_0))=k2_struct_0(esk1_0)|k2_struct_0(esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_35]), c_0_20])])).
cnf(c_0_53, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,esk2_0),k2_struct_0(esk1_0))=k2_struct_0(k6_tmap_1(esk1_0,esk2_0))|k2_struct_0(esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_35]), c_0_20])])).
cnf(c_0_54, plain, (k8_relset_1(X2,X2,k6_partfun1(X2),X1)=X1|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_55, plain, (v2_struct_0(X1)|k7_tmap_1(X1,X2)=k6_partfun1(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_56, negated_conjecture, (k10_relat_1(X1,k2_struct_0(X2))=u1_struct_0(esk1_0)|u1_struct_0(esk1_0)!=k1_xboole_0|~l1_struct_0(X2)|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X2))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_26]), c_0_36])])).
cnf(c_0_57, negated_conjecture, (k8_relset_1(u1_struct_0(X1),u1_struct_0(esk1_0),X2,k2_struct_0(k6_tmap_1(esk1_0,esk2_0)))=k2_struct_0(X1)|k2_struct_0(X1)!=k1_xboole_0|~l1_struct_0(k6_tmap_1(esk1_0,esk2_0))|~l1_struct_0(X1)|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))), inference(spm,[status(thm)],[c_0_38, c_0_26])).
cnf(c_0_58, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1,k2_struct_0(k6_tmap_1(esk1_0,esk2_0)))=k2_struct_0(k6_tmap_1(esk1_0,esk2_0))|k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_35]), c_0_36])])).
cnf(c_0_59, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0))=u1_struct_0(esk1_0)|u1_struct_0(esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_45]), c_0_20])])).
cnf(c_0_60, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0))=k2_struct_0(k6_tmap_1(esk1_0,esk2_0))|u1_struct_0(esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_45]), c_0_20])])).
cnf(c_0_61, plain, (k10_relat_1(k6_partfun1(X1),X2)=X2|~m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_54, c_0_37])).
cnf(c_0_62, negated_conjecture, (k6_partfun1(u1_struct_0(esk1_0))=k7_tmap_1(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_19]), c_0_20]), c_0_21])]), c_0_22])).
fof(c_0_63, plain, ![X26, X27, X28]:(v2_struct_0(X26)|(~v2_pre_topc(X26)|~l1_pre_topc(X26))|(~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X26)))|(~m1_subset_1(X28,k1_zfmisc_1(u1_struct_0(k6_tmap_1(X26,X27))))|(X28!=X27|v3_pre_topc(X28,k6_tmap_1(X26,X27)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t105_tmap_1])])])])).
cnf(c_0_64, negated_conjecture, (k8_relset_1(X1,X2,X3,k2_struct_0(X4))=u1_struct_0(esk1_0)|u1_struct_0(esk1_0)!=k1_xboole_0|~l1_struct_0(X4)|~v1_funct_2(X3,u1_struct_0(esk1_0),u1_struct_0(X4))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X4))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_37, c_0_56])).
cnf(c_0_65, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1,k2_struct_0(k6_tmap_1(esk1_0,esk2_0)))=k2_struct_0(k6_tmap_1(esk1_0,esk2_0))|~l1_struct_0(k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_26]), c_0_58])).
cnf(c_0_66, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=u1_struct_0(esk1_0)|u1_struct_0(esk1_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_59, c_0_60])).
fof(c_0_67, plain, ![X29, X30, X31, X32]:(((~v5_pre_topc(X31,X29,X30)|(~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X30)))|(~v3_pre_topc(X32,X30)|v3_pre_topc(k8_relset_1(u1_struct_0(X29),u1_struct_0(X30),X31,X32),X29)))|k2_struct_0(X30)=k1_xboole_0|(~v1_funct_1(X31)|(~v1_funct_2(X31,u1_struct_0(X29),u1_struct_0(X30))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X29),u1_struct_0(X30))))))|~l1_pre_topc(X30)|~l1_pre_topc(X29))&((m1_subset_1(esk3_3(X29,X30,X31),k1_zfmisc_1(u1_struct_0(X30)))|v5_pre_topc(X31,X29,X30)|k2_struct_0(X30)=k1_xboole_0|(~v1_funct_1(X31)|(~v1_funct_2(X31,u1_struct_0(X29),u1_struct_0(X30))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X29),u1_struct_0(X30))))))|~l1_pre_topc(X30)|~l1_pre_topc(X29))&((v3_pre_topc(esk3_3(X29,X30,X31),X30)|v5_pre_topc(X31,X29,X30)|k2_struct_0(X30)=k1_xboole_0|(~v1_funct_1(X31)|(~v1_funct_2(X31,u1_struct_0(X29),u1_struct_0(X30))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X29),u1_struct_0(X30))))))|~l1_pre_topc(X30)|~l1_pre_topc(X29))&(~v3_pre_topc(k8_relset_1(u1_struct_0(X29),u1_struct_0(X30),X31,esk3_3(X29,X30,X31)),X29)|v5_pre_topc(X31,X29,X30)|k2_struct_0(X30)=k1_xboole_0|(~v1_funct_1(X31)|(~v1_funct_2(X31,u1_struct_0(X29),u1_struct_0(X30))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X29),u1_struct_0(X30))))))|~l1_pre_topc(X30)|~l1_pre_topc(X29)))))&((~v5_pre_topc(X31,X29,X30)|(~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X30)))|(~v3_pre_topc(X32,X30)|v3_pre_topc(k8_relset_1(u1_struct_0(X29),u1_struct_0(X30),X31,X32),X29)))|k2_struct_0(X29)!=k1_xboole_0|(~v1_funct_1(X31)|(~v1_funct_2(X31,u1_struct_0(X29),u1_struct_0(X30))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X29),u1_struct_0(X30))))))|~l1_pre_topc(X30)|~l1_pre_topc(X29))&((m1_subset_1(esk3_3(X29,X30,X31),k1_zfmisc_1(u1_struct_0(X30)))|v5_pre_topc(X31,X29,X30)|k2_struct_0(X29)!=k1_xboole_0|(~v1_funct_1(X31)|(~v1_funct_2(X31,u1_struct_0(X29),u1_struct_0(X30))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X29),u1_struct_0(X30))))))|~l1_pre_topc(X30)|~l1_pre_topc(X29))&((v3_pre_topc(esk3_3(X29,X30,X31),X30)|v5_pre_topc(X31,X29,X30)|k2_struct_0(X29)!=k1_xboole_0|(~v1_funct_1(X31)|(~v1_funct_2(X31,u1_struct_0(X29),u1_struct_0(X30))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X29),u1_struct_0(X30))))))|~l1_pre_topc(X30)|~l1_pre_topc(X29))&(~v3_pre_topc(k8_relset_1(u1_struct_0(X29),u1_struct_0(X30),X31,esk3_3(X29,X30,X31)),X29)|v5_pre_topc(X31,X29,X30)|k2_struct_0(X29)!=k1_xboole_0|(~v1_funct_1(X31)|(~v1_funct_2(X31,u1_struct_0(X29),u1_struct_0(X30))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X29),u1_struct_0(X30))))))|~l1_pre_topc(X30)|~l1_pre_topc(X29)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_tops_2])])])])])).
cnf(c_0_68, negated_conjecture, (k10_relat_1(k7_tmap_1(esk1_0,esk2_0),esk2_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_19]), c_0_62]), c_0_62]), c_0_42])])).
cnf(c_0_69, plain, (v2_struct_0(X1)|v3_pre_topc(X3,k6_tmap_1(X1,X2))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k6_tmap_1(X1,X2))))|X3!=X2), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_70, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=u1_struct_0(esk1_0)|~l1_struct_0(k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_26]), c_0_26])]), c_0_66])).
cnf(c_0_71, plain, (v3_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(X3),X1,X4),X2)|~v5_pre_topc(X1,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~v3_pre_topc(X4,X3)|k2_struct_0(X2)!=k1_xboole_0|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(X3)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_72, negated_conjecture, (k8_relset_1(X1,X2,k7_tmap_1(esk1_0,esk2_0),esk2_0)=esk2_0|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_37, c_0_68])).
cnf(c_0_73, plain, (v3_pre_topc(X1,k6_tmap_1(X2,X1))|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k6_tmap_1(X2,X1))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(er,[status(thm)],[c_0_69])).
cnf(c_0_74, plain, (v5_pre_topc(X3,X1,X2)|~v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk3_3(X1,X2,X3)),X1)|k2_struct_0(X1)!=k1_xboole_0|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_75, plain, (v3_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(X3),X1,X4),X2)|k2_struct_0(X3)=k1_xboole_0|~v5_pre_topc(X1,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~v3_pre_topc(X4,X3)|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(X3)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_76, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=u1_struct_0(esk1_0)|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_35]), c_0_36])])).
cnf(c_0_77, negated_conjecture, (v3_pre_topc(esk2_0,X1)|k2_struct_0(X1)!=k1_xboole_0|~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),X1,X2)|~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(X1),u1_struct_0(X2))|~v3_pre_topc(esk2_0,X2)|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_41])])).
cnf(c_0_78, negated_conjecture, (v3_pre_topc(esk2_0,k6_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_19]), c_0_26]), c_0_19]), c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_79, plain, (v5_pre_topc(X3,X1,X2)|k2_struct_0(X2)=k1_xboole_0|~v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk3_3(X1,X2,X3)),X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_80, negated_conjecture, (v5_pre_topc(X1,X2,k6_tmap_1(esk1_0,esk2_0))|k2_struct_0(X2)!=k1_xboole_0|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~v3_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(esk1_0),X1,esk3_3(X2,k6_tmap_1(esk1_0,esk2_0),X1)),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_26]), c_0_36])])).
cnf(c_0_81, plain, (v3_pre_topc(esk3_3(X1,X2,X3),X2)|v5_pre_topc(X3,X1,X2)|k2_struct_0(X2)=k1_xboole_0|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_82, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(esk1_0),X2,X3),X1)|~v5_pre_topc(X2,X1,k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2)|~v3_pre_topc(X3,k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_26]), c_0_36])])).
cnf(c_0_83, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_40]), c_0_41]), c_0_42])])).
cnf(c_0_84, negated_conjecture, (v3_pre_topc(esk2_0,X1)|k2_struct_0(X1)!=k1_xboole_0|~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),X1,k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(X1),u1_struct_0(esk1_0))|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_26]), c_0_78]), c_0_19]), c_0_36])])).
cnf(c_0_85, negated_conjecture, (v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|v3_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_86, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v5_pre_topc(X1,X2,k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~v3_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(esk1_0),X1,esk3_3(X2,k6_tmap_1(esk1_0,esk2_0),X1)),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_26]), c_0_36])])).
cnf(c_0_87, negated_conjecture, (v5_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0))|k2_struct_0(k6_tmap_1(esk1_0,esk2_0))!=k1_xboole_0|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~v3_pre_topc(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1,esk3_3(k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),X1)),k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_26]), c_0_36])])).
cnf(c_0_88, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v5_pre_topc(X1,X2,k6_tmap_1(esk1_0,esk2_0))|v3_pre_topc(esk3_3(X2,k6_tmap_1(esk1_0,esk2_0),X1),k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_26]), c_0_36])])).
cnf(c_0_89, negated_conjecture, (~v3_pre_topc(esk2_0,esk1_0)|~v1_funct_1(k7_tmap_1(esk1_0,esk2_0))|~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))|~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_90, negated_conjecture, (k8_relset_1(u1_struct_0(X1),u1_struct_0(esk1_0),X2,k2_struct_0(k6_tmap_1(esk1_0,esk2_0)))=k2_struct_0(X1)|k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|~l1_struct_0(k6_tmap_1(esk1_0,esk2_0))|~l1_struct_0(X1)|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_91, plain, ![X34, X35, X36, X37]:((~v5_pre_topc(X36,X34,X35)|v5_pre_topc(X37,g1_pre_topc(u1_struct_0(X34),u1_pre_topc(X34)),X35)|X36!=X37|(~v1_funct_1(X37)|(~v1_funct_2(X37,u1_struct_0(g1_pre_topc(u1_struct_0(X34),u1_pre_topc(X34))),u1_struct_0(X35))|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X34),u1_pre_topc(X34))),u1_struct_0(X35))))))|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(X34),u1_struct_0(X35))|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X34),u1_struct_0(X35))))))|(~v2_pre_topc(X35)|~l1_pre_topc(X35))|(~v2_pre_topc(X34)|~l1_pre_topc(X34)))&(~v5_pre_topc(X37,g1_pre_topc(u1_struct_0(X34),u1_pre_topc(X34)),X35)|v5_pre_topc(X36,X34,X35)|X36!=X37|(~v1_funct_1(X37)|(~v1_funct_2(X37,u1_struct_0(g1_pre_topc(u1_struct_0(X34),u1_pre_topc(X34))),u1_struct_0(X35))|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X34),u1_pre_topc(X34))),u1_struct_0(X35))))))|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(X34),u1_struct_0(X35))|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X34),u1_struct_0(X35))))))|(~v2_pre_topc(X35)|~l1_pre_topc(X35))|(~v2_pre_topc(X34)|~l1_pre_topc(X34)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t62_pre_topc])])])])).
fof(c_0_92, plain, ![X38, X39]:((~v3_pre_topc(X39,X38)|g1_pre_topc(u1_struct_0(X38),u1_pre_topc(X38))=k6_tmap_1(X38,X39)|~m1_subset_1(X39,k1_zfmisc_1(u1_struct_0(X38)))|(v2_struct_0(X38)|(~v2_pre_topc(X38)|~l1_pre_topc(X38))))&(g1_pre_topc(u1_struct_0(X38),u1_pre_topc(X38))!=k6_tmap_1(X38,X39)|v3_pre_topc(X39,X38)|~m1_subset_1(X39,k1_zfmisc_1(u1_struct_0(X38)))|(v2_struct_0(X38)|(~v2_pre_topc(X38)|~l1_pre_topc(X38))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t106_tmap_1])])])])])).
cnf(c_0_93, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(esk1_0),X2,X3),X1)|~v5_pre_topc(X2,X1,k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2)|~v3_pre_topc(X3,k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(X1)), inference(rw,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_94, negated_conjecture, (v3_pre_topc(esk2_0,esk1_0)|k2_struct_0(esk1_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_40]), c_0_42]), c_0_20])])).
cnf(c_0_95, negated_conjecture, (v5_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~v3_pre_topc(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1,esk3_3(k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),X1)),k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_26]), c_0_36])]), c_0_87])).
cnf(c_0_96, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v5_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0))|v3_pre_topc(esk3_3(k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),X1),k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_26]), c_0_36])])).
cnf(c_0_97, negated_conjecture, (~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))|~v3_pre_topc(esk2_0,esk1_0)|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk2_0)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_41])])).
cnf(c_0_98, negated_conjecture, (k8_relset_1(u1_struct_0(X1),u1_struct_0(esk1_0),X2,k2_struct_0(k6_tmap_1(esk1_0,esk2_0)))=k2_struct_0(X1)|k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|~l1_struct_0(X1)|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_35]), c_0_36])])).
cnf(c_0_99, plain, (v5_pre_topc(X4,X2,X3)|~v5_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)),X3)|X4!=X1|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))),u1_struct_0(X3))))|~v1_funct_1(X4)|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_100, plain, (g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))=k6_tmap_1(X2,X1)|v2_struct_0(X2)|~v3_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_101, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(esk1_0),X2,esk2_0),X1)|~v5_pre_topc(X2,X1,k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_19]), c_0_78])])).
cnf(c_0_102, negated_conjecture, (v3_pre_topc(esk2_0,esk1_0)|u1_struct_0(esk1_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_45]), c_0_20])])).
cnf(c_0_103, plain, (m1_subset_1(esk3_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|v5_pre_topc(X3,X1,X2)|k2_struct_0(X2)=k1_xboole_0|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_104, negated_conjecture, (v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0))|~v3_pre_topc(esk3_3(k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),k7_tmap_1(esk1_0,esk2_0)),k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(esk3_3(k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),k7_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_54]), c_0_62]), c_0_62]), c_0_40]), c_0_62]), c_0_41]), c_0_62]), c_0_62]), c_0_42]), c_0_62])])).
cnf(c_0_105, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0))|v3_pre_topc(esk3_3(k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),k7_tmap_1(esk1_0,esk2_0)),k6_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_40]), c_0_41]), c_0_42])])).
cnf(c_0_106, negated_conjecture, (~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v3_pre_topc(esk2_0,esk1_0)|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_26]), c_0_26])).
cnf(c_0_107, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,esk2_0),k2_struct_0(k6_tmap_1(esk1_0,esk2_0)))=k2_struct_0(esk1_0)|k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|~l1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_40]), c_0_41]), c_0_42])])).
cnf(c_0_108, plain, (v5_pre_topc(X1,X2,X3)|~v5_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)),X3)|~v1_funct_2(X1,u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))),u1_struct_0(X3))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~v2_pre_topc(X2)), inference(er,[status(thm)],[c_0_99])).
cnf(c_0_109, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))=k6_tmap_1(esk1_0,esk2_0)|~v3_pre_topc(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_19]), c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_110, negated_conjecture, (v3_pre_topc(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,esk2_0),esk2_0),esk1_0)|v3_pre_topc(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_85]), c_0_40]), c_0_41]), c_0_42]), c_0_20])]), c_0_102])).
cnf(c_0_111, negated_conjecture, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(X2,k6_tmap_1(esk1_0,esk2_0),X1)|m1_subset_1(esk3_3(k6_tmap_1(esk1_0,esk2_0),X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X2,u1_struct_0(esk1_0),u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_26]), c_0_36])])).
cnf(c_0_112, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(esk3_3(k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),k7_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_113, negated_conjecture, (~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~v3_pre_topc(esk2_0,esk1_0)|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_106, c_0_40])])).
cnf(c_0_114, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,esk2_0),k2_struct_0(k6_tmap_1(esk1_0,esk2_0)))=k2_struct_0(k6_tmap_1(esk1_0,esk2_0))|k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_40]), c_0_41]), c_0_42])])).
cnf(c_0_115, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,esk2_0),k2_struct_0(k6_tmap_1(esk1_0,esk2_0)))=k2_struct_0(esk1_0)|k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_35]), c_0_20])])).
cnf(c_0_116, negated_conjecture, (v5_pre_topc(X1,esk1_0,X2)|~v5_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0),X2)|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(X2))|~v1_funct_1(X1)|~v3_pre_topc(esk2_0,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X2))))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_26]), c_0_26]), c_0_20]), c_0_21])])).
cnf(c_0_117, negated_conjecture, (v3_pre_topc(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_37]), c_0_68]), c_0_42])])).
cnf(c_0_118, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|v5_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0))|m1_subset_1(esk3_3(k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_26]), c_0_36])]), c_0_83])).
cnf(c_0_119, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0))|~m1_subset_1(esk3_3(k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),k7_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(rw,[status(thm)],[c_0_112, c_0_83])).
cnf(c_0_120, plain, (v2_pre_topc(k6_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_121, negated_conjecture, (~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))|~v3_pre_topc(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_113, c_0_42])])).
cnf(c_0_122, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,esk2_0),X1)=X1|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_54, c_0_62])).
cnf(c_0_123, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k2_struct_0(esk1_0)|k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_124, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k2_struct_0(esk1_0)|k2_struct_0(esk1_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_125, negated_conjecture, (v5_pre_topc(X1,esk1_0,X2)|~v5_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0),X2)|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X2))))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_116, c_0_117])])).
cnf(c_0_126, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_40]), c_0_41]), c_0_42])]), c_0_119])).
cnf(c_0_127, negated_conjecture, (v2_pre_topc(k6_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_19]), c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_128, negated_conjecture, (~v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),esk1_0,k6_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_121, c_0_117])])).
cnf(c_0_129, plain, (m1_subset_1(esk3_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|v5_pre_topc(X3,X1,X2)|k2_struct_0(X1)!=k1_xboole_0|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_130, negated_conjecture, (k10_relat_1(k7_tmap_1(esk1_0,esk2_0),X1)=X1|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_37, c_0_122])).
cnf(c_0_131, plain, (v3_pre_topc(esk3_3(X1,X2,X3),X2)|v5_pre_topc(X3,X1,X2)|k2_struct_0(X1)!=k1_xboole_0|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_132, negated_conjecture, (k2_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0|k2_struct_0(esk1_0)!=k1_xboole_0), inference(ef,[status(thm)],[c_0_123])).
cnf(c_0_133, negated_conjecture, (k2_struct_0(esk1_0)=u1_struct_0(esk1_0)|k2_struct_0(esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_124]), c_0_26]), c_0_36])])).
cnf(c_0_134, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_26]), c_0_40]), c_0_41]), c_0_26]), c_0_42]), c_0_36]), c_0_127])]), c_0_128])).
cnf(c_0_135, plain, (v5_pre_topc(X1,X2,X3)|m1_subset_1(esk3_3(X2,X3,X1),k1_zfmisc_1(u1_struct_0(X3)))|u1_struct_0(X2)!=k1_xboole_0|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(X3)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_129, c_0_45])).
cnf(c_0_136, negated_conjecture, (k10_relat_1(k7_tmap_1(esk1_0,esk2_0),X1)=X1|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_130, c_0_42])])).
cnf(c_0_137, negated_conjecture, (v5_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0),X2)|v3_pre_topc(esk3_3(k6_tmap_1(esk1_0,esk2_0),X2,X1),X2)|k2_struct_0(esk1_0)!=k1_xboole_0|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X2))))|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_132]), c_0_26]), c_0_26]), c_0_36])])).
cnf(c_0_138, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_133, c_0_134])])).
cnf(c_0_139, negated_conjecture, (v5_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0),X2)|m1_subset_1(esk3_3(k6_tmap_1(esk1_0,esk2_0),X2,X1),k1_zfmisc_1(u1_struct_0(X2)))|u1_struct_0(esk1_0)!=k1_xboole_0|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X2))))|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_26]), c_0_36])])).
cnf(c_0_140, negated_conjecture, (k8_relset_1(X1,X2,k7_tmap_1(esk1_0,esk2_0),X3)=X3|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_37, c_0_136])).
cnf(c_0_141, negated_conjecture, (v5_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0),X2)|v3_pre_topc(esk3_3(k6_tmap_1(esk1_0,esk2_0),X2,X1),X2)|~v1_funct_2(X1,k1_xboole_0,u1_struct_0(X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,u1_struct_0(X2))))|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_137, c_0_134]), c_0_134]), c_0_138])])).
cnf(c_0_142, negated_conjecture, (u1_struct_0(k6_tmap_1(esk1_0,esk2_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_26, c_0_134])).
cnf(c_0_143, negated_conjecture, (v5_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0),X2)|m1_subset_1(esk3_3(k6_tmap_1(esk1_0,esk2_0),X2,X1),k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X1,k1_xboole_0,u1_struct_0(X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,u1_struct_0(X2))))|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_139, c_0_134]), c_0_134]), c_0_134])])).
cnf(c_0_144, negated_conjecture, (v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),X1,X2)|k2_struct_0(X1)!=k1_xboole_0|~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(X1),u1_struct_0(X2))|~v3_pre_topc(esk3_3(X1,X2,k7_tmap_1(esk1_0,esk2_0)),X1)|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk3_3(X1,X2,k7_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_140]), c_0_41])])).
cnf(c_0_145, negated_conjecture, (v5_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0))|v3_pre_topc(esk3_3(k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),X1),k6_tmap_1(esk1_0,esk2_0))|~v1_funct_2(X1,k1_xboole_0,k1_xboole_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_142]), c_0_36])])).
cnf(c_0_146, negated_conjecture, (v1_funct_2(k7_tmap_1(esk1_0,esk2_0),k1_xboole_0,k1_xboole_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_134]), c_0_134])).
cnf(c_0_147, negated_conjecture, (m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_134]), c_0_134])).
cnf(c_0_148, negated_conjecture, (v5_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0))|m1_subset_1(esk3_3(k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),X1),k1_zfmisc_1(k1_xboole_0))|~v1_funct_2(X1,k1_xboole_0,k1_xboole_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_142]), c_0_36])])).
cnf(c_0_149, negated_conjecture, (v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),X1,X2)|k2_struct_0(X1)!=k1_xboole_0|~v1_funct_2(k7_tmap_1(esk1_0,esk2_0),u1_struct_0(X1),u1_struct_0(X2))|~v3_pre_topc(esk3_3(X1,X2,k7_tmap_1(esk1_0,esk2_0)),X1)|~m1_subset_1(k7_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk3_3(X1,X2,k7_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(k1_xboole_0))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(rw,[status(thm)],[c_0_144, c_0_134])).
cnf(c_0_150, negated_conjecture, (v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0))|v3_pre_topc(esk3_3(k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),k7_tmap_1(esk1_0,esk2_0)),k6_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146]), c_0_41]), c_0_147])])).
cnf(c_0_151, negated_conjecture, (v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0))|m1_subset_1(esk3_3(k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),k7_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_146]), c_0_41]), c_0_147])])).
cnf(c_0_152, negated_conjecture, (v5_pre_topc(X1,esk1_0,X2)|~v5_pre_topc(X1,k6_tmap_1(esk1_0,esk2_0),X2)|~v1_funct_2(X1,k1_xboole_0,u1_struct_0(X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,u1_struct_0(X2))))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_125, c_0_134]), c_0_134])).
cnf(c_0_153, negated_conjecture, (v5_pre_topc(k7_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0),k6_tmap_1(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_150]), c_0_83]), c_0_134]), c_0_26]), c_0_134]), c_0_26]), c_0_134]), c_0_146]), c_0_26]), c_0_134]), c_0_26]), c_0_134]), c_0_147]), c_0_36])]), c_0_151])).
cnf(c_0_154, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_153]), c_0_26]), c_0_134]), c_0_146]), c_0_41]), c_0_26]), c_0_134]), c_0_147]), c_0_36]), c_0_127])]), c_0_128]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 155
# Proof object clause steps            : 126
# Proof object formula steps           : 29
# Proof object conjectures             : 98
# Proof object clause conjectures      : 95
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 30
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 77
# Proof object simplifying inferences  : 242
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 37
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 37
# Processed clauses                    : 12614
# ...of these trivial                  : 241
# ...subsumed                          : 9139
# ...remaining for further processing  : 3234
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 302
# Backward-rewritten                   : 873
# Generated clauses                    : 34700
# ...of the previous two non-trivial   : 32603
# Contextual simplify-reflections      : 1097
# Paramodulations                      : 34628
# Factorizations                       : 3
# Equation resolutions                 : 41
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1991
#    Positive orientable unit clauses  : 33
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 1954
# Current number of unprocessed clauses: 16397
# ...number of literals in the above   : 132840
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1240
# Clause-clause subsumption calls (NU) : 2063462
# Rec. Clause-clause subsumption calls : 220944
# Non-unit clause-clause subsumptions  : 10276
# Unit Clause-clause subsumption calls : 7168
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 112
# BW rewrite match successes           : 20
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2175392

# -------------------------------------------------
# User time                : 1.990 s
# System time              : 0.013 s
# Total time               : 2.003 s
# Maximum resident set size: 28776 pages
