# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t16_funct_1, conjecture, ![X1]:(![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X2)=X1&k1_relat_1(X3)=X1)=>X2=X3)))=>X1=k1_xboole_0), file('data/MPTP/proofs/140769b26d06467c86260e9e5595c9df.E_input', t16_funct_1)).
fof(s3_funct_1__e7_25__funct_1, axiom, ![X1]:?[X2]:((v1_relat_1(X2)&v1_funct_1(X2))&(k1_relat_1(X2)=X1&![X3]:(r2_hidden(X3,X1)=>k1_funct_1(X2,X3)=1))), file('data/MPTP/proofs/140769b26d06467c86260e9e5595c9df.E_input', s3_funct_1__e7_25__funct_1)).
fof(t11_setfam_1, axiom, ![X1]:k1_setfam_1(k1_tarski(X1))=X1, file('data/MPTP/proofs/140769b26d06467c86260e9e5595c9df.E_input', t11_setfam_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/140769b26d06467c86260e9e5595c9df.E_input', t69_enumset1)).
fof(t15_funct_1, axiom, ![X1]:(~(X1=k1_xboole_0)=>![X2]:?[X3]:((v1_relat_1(X3)&v1_funct_1(X3))&(k1_relat_1(X3)=X1&k2_relat_1(X3)=k1_tarski(X2)))), file('data/MPTP/proofs/140769b26d06467c86260e9e5595c9df.E_input', t15_funct_1)).
fof(c_0_5, negated_conjecture, ~(![X1]:(![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X2)=X1&k1_relat_1(X3)=X1)=>X2=X3)))=>X1=k1_xboole_0)), inference(assume_negation,[status(cth)],[t16_funct_1])).
fof(c_0_6, negated_conjecture, ![X5, X6]:((~v1_relat_1(X5)|~v1_funct_1(X5)|(~v1_relat_1(X6)|~v1_funct_1(X6)|(k1_relat_1(X5)!=esk1_0|k1_relat_1(X6)!=esk1_0|X5=X6)))&esk1_0!=k1_xboole_0), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
fof(c_0_7, plain, ![X7, X9]:((v1_relat_1(esk2_1(X7))&v1_funct_1(esk2_1(X7)))&(k1_relat_1(esk2_1(X7))=X7&(~r2_hidden(X9,X7)|k1_funct_1(esk2_1(X7),X9)=1))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[s3_funct_1__e7_25__funct_1])])])])).
fof(c_0_8, plain, ![X13]:k1_setfam_1(k1_tarski(X13))=X13, inference(variable_rename,[status(thm)],[t11_setfam_1])).
fof(c_0_9, plain, ![X14]:k2_tarski(X14,X14)=k1_tarski(X14), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_10, plain, ![X10, X11]:(((v1_relat_1(esk3_2(X10,X11))|X10=k1_xboole_0)&(v1_funct_1(esk3_2(X10,X11))|X10=k1_xboole_0))&((k1_relat_1(esk3_2(X10,X11))=X10|X10=k1_xboole_0)&(k2_relat_1(esk3_2(X10,X11))=k1_tarski(X11)|X10=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t15_funct_1])])])])])])).
cnf(c_0_11, negated_conjecture, (X1=X2|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|k1_relat_1(X1)!=esk1_0|k1_relat_1(X2)!=esk1_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12, plain, (v1_funct_1(esk2_1(X1))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, plain, (k1_relat_1(esk2_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14, plain, (v1_relat_1(esk2_1(X1))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_15, plain, (k1_setfam_1(k1_tarski(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, plain, (k2_relat_1(esk3_2(X1,X2))=k1_tarski(X2)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, negated_conjecture, (X1=esk2_1(esk1_0)|k1_relat_1(X1)!=esk1_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13]), c_0_14])])])).
cnf(c_0_19, plain, (v1_funct_1(esk3_2(X1,X2))|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_20, plain, (v1_relat_1(esk3_2(X1,X2))|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_21, plain, (k1_setfam_1(k2_tarski(X1,X1))=X1), inference(rw,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_22, plain, (X1=k1_xboole_0|k2_relat_1(esk3_2(X1,X2))=k2_tarski(X2,X2)), inference(rw,[status(thm)],[c_0_17, c_0_16])).
cnf(c_0_23, negated_conjecture, (esk3_2(X1,X2)=esk2_1(esk1_0)|X1=k1_xboole_0|k1_relat_1(esk3_2(X1,X2))!=esk1_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])).
cnf(c_0_24, plain, (k1_relat_1(esk3_2(X1,X2))=X1|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_25, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_26, plain, (k1_setfam_1(k2_relat_1(esk3_2(X1,X2)))=X2|X1=k1_xboole_0), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27, negated_conjecture, (esk3_2(esk1_0,X1)=esk2_1(esk1_0)), inference(sr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24])]), c_0_25])).
cnf(c_0_28, negated_conjecture, (k1_setfam_1(k2_relat_1(esk2_1(esk1_0)))=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_25])).
cnf(c_0_29, negated_conjecture, (X1=X2), inference(spm,[status(thm)],[c_0_28, c_0_28])).
cnf(c_0_30, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_25, c_0_29]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 31
# Proof object clause steps            : 20
# Proof object formula steps           : 11
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 12
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 11
# Processed clauses                    : 29
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 28
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 12
# Backward-rewritten                   : 1
# Generated clauses                    : 41
# ...of the previous two non-trivial   : 29
# Contextual simplify-reflections      : 1
# Paramodulations                      : 37
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 3
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 22
# ...number of literals in the above   : 42
# Current number of archived formulas  : 0
# Current number of archived clauses   : 26
# Clause-clause subsumption calls (NU) : 2
# Rec. Clause-clause subsumption calls : 1
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 48
# BW rewrite match attempts            : 53
# BW rewrite match successes           : 53
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 940

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.003 s
# Total time               : 0.009 s
# Maximum resident set size: 28792 pages
