# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(r2_wellord1(X1,X2)<=>(r1_relat_2(X1,X2)&(r8_relat_2(X1,X2)&(r4_relat_2(X1,X2)&(r6_relat_2(X1,X2)&r1_wellord1(X1,X2))))))), file('data/MPTP/proofs/14f621855177424ba1aedeb027529ad4.E_input', d5_wellord1)).
fof(t5_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v1_wellord1(X1)<=>r1_wellord1(X1,k3_relat_1(X1)))), file('data/MPTP/proofs/14f621855177424ba1aedeb027529ad4.E_input', t5_wellord1)).
fof(d14_relat_2, axiom, ![X1]:(v1_relat_1(X1)=>(v6_relat_2(X1)<=>r6_relat_2(X1,k3_relat_1(X1)))), file('data/MPTP/proofs/14f621855177424ba1aedeb027529ad4.E_input', d14_relat_2)).
fof(d16_relat_2, axiom, ![X1]:(v1_relat_1(X1)=>(v8_relat_2(X1)<=>r8_relat_2(X1,k3_relat_1(X1)))), file('data/MPTP/proofs/14f621855177424ba1aedeb027529ad4.E_input', d16_relat_2)).
fof(t53_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r3_wellord1(X1,X2,X3)=>((v1_relat_2(X1)=>v1_relat_2(X2))&((v8_relat_2(X1)=>v8_relat_2(X2))&((v6_relat_2(X1)=>v6_relat_2(X2))&((v4_relat_2(X1)=>v4_relat_2(X2))&(v1_wellord1(X1)=>v1_wellord1(X2)))))))))), file('data/MPTP/proofs/14f621855177424ba1aedeb027529ad4.E_input', t53_wellord1)).
fof(t54_wellord1, conjecture, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((v2_wellord1(X1)&r3_wellord1(X1,X2,X3))=>v2_wellord1(X2))))), file('data/MPTP/proofs/14f621855177424ba1aedeb027529ad4.E_input', t54_wellord1)).
fof(d9_relat_2, axiom, ![X1]:(v1_relat_1(X1)=>(v1_relat_2(X1)<=>r1_relat_2(X1,k3_relat_1(X1)))), file('data/MPTP/proofs/14f621855177424ba1aedeb027529ad4.E_input', d9_relat_2)).
fof(d12_relat_2, axiom, ![X1]:(v1_relat_1(X1)=>(v4_relat_2(X1)<=>r4_relat_2(X1,k3_relat_1(X1)))), file('data/MPTP/proofs/14f621855177424ba1aedeb027529ad4.E_input', d12_relat_2)).
fof(t8_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(r2_wellord1(X1,k3_relat_1(X1))<=>v2_wellord1(X1))), file('data/MPTP/proofs/14f621855177424ba1aedeb027529ad4.E_input', t8_wellord1)).
fof(c_0_9, plain, ![X12, X13]:(((r1_relat_2(X12,X13)|~r2_wellord1(X12,X13)|~v1_relat_1(X12))&((r8_relat_2(X12,X13)|~r2_wellord1(X12,X13)|~v1_relat_1(X12))&((r4_relat_2(X12,X13)|~r2_wellord1(X12,X13)|~v1_relat_1(X12))&((r6_relat_2(X12,X13)|~r2_wellord1(X12,X13)|~v1_relat_1(X12))&(r1_wellord1(X12,X13)|~r2_wellord1(X12,X13)|~v1_relat_1(X12))))))&(~r1_relat_2(X12,X13)|(~r8_relat_2(X12,X13)|(~r4_relat_2(X12,X13)|(~r6_relat_2(X12,X13)|~r1_wellord1(X12,X13))))|r2_wellord1(X12,X13)|~v1_relat_1(X12))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_wellord1])])])])).
fof(c_0_10, plain, ![X16]:((~v1_wellord1(X16)|r1_wellord1(X16,k3_relat_1(X16))|~v1_relat_1(X16))&(~r1_wellord1(X16,k3_relat_1(X16))|v1_wellord1(X16)|~v1_relat_1(X16))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_wellord1])])])).
cnf(c_0_11, plain, (r2_wellord1(X1,X2)|~r1_relat_2(X1,X2)|~r8_relat_2(X1,X2)|~r4_relat_2(X1,X2)|~r6_relat_2(X1,X2)|~r1_wellord1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12, plain, (r1_wellord1(X1,k3_relat_1(X1))|~v1_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_13, plain, ![X14]:((~v6_relat_2(X14)|r6_relat_2(X14,k3_relat_1(X14))|~v1_relat_1(X14))&(~r6_relat_2(X14,k3_relat_1(X14))|v6_relat_2(X14)|~v1_relat_1(X14))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d14_relat_2])])])).
fof(c_0_14, plain, ![X1, X2]:(epred1_2(X2,X1)<=>((v1_relat_2(X1)=>v1_relat_2(X2))&((v8_relat_2(X1)=>v8_relat_2(X2))&((v6_relat_2(X1)=>v6_relat_2(X2))&((v4_relat_2(X1)=>v4_relat_2(X2))&(v1_wellord1(X1)=>v1_wellord1(X2))))))), introduced(definition)).
cnf(c_0_15, plain, (r2_wellord1(X1,k3_relat_1(X1))|~r6_relat_2(X1,k3_relat_1(X1))|~r8_relat_2(X1,k3_relat_1(X1))|~v1_wellord1(X1)|~r1_relat_2(X1,k3_relat_1(X1))|~r4_relat_2(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16, plain, (r6_relat_2(X1,k3_relat_1(X1))|~v6_relat_2(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_17, plain, ![X15]:((~v8_relat_2(X15)|r8_relat_2(X15,k3_relat_1(X15))|~v1_relat_1(X15))&(~r8_relat_2(X15,k3_relat_1(X15))|v8_relat_2(X15)|~v1_relat_1(X15))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d16_relat_2])])])).
fof(c_0_18, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r3_wellord1(X1,X2,X3)=>epred1_2(X2,X1))))), inference(apply_def,[status(thm)],[t53_wellord1, c_0_14])).
fof(c_0_19, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((v2_wellord1(X1)&r3_wellord1(X1,X2,X3))=>v2_wellord1(X2)))))), inference(assume_negation,[status(cth)],[t54_wellord1])).
cnf(c_0_20, plain, (r2_wellord1(X1,k3_relat_1(X1))|~r8_relat_2(X1,k3_relat_1(X1))|~v1_wellord1(X1)|~v6_relat_2(X1)|~r1_relat_2(X1,k3_relat_1(X1))|~r4_relat_2(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_21, plain, (r8_relat_2(X1,k3_relat_1(X1))|~v8_relat_2(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_22, plain, ![X8]:((~v1_relat_2(X8)|r1_relat_2(X8,k3_relat_1(X8))|~v1_relat_1(X8))&(~r1_relat_2(X8,k3_relat_1(X8))|v1_relat_2(X8)|~v1_relat_1(X8))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_relat_2])])])).
fof(c_0_23, plain, ![X1, X2]:(epred1_2(X2,X1)=>((v1_relat_2(X1)=>v1_relat_2(X2))&((v8_relat_2(X1)=>v8_relat_2(X2))&((v6_relat_2(X1)=>v6_relat_2(X2))&((v4_relat_2(X1)=>v4_relat_2(X2))&(v1_wellord1(X1)=>v1_wellord1(X2))))))), inference(split_equiv,[status(thm)],[c_0_14])).
fof(c_0_24, plain, ![X9, X10, X11]:(~v1_relat_1(X9)|(~v1_relat_1(X10)|(~v1_relat_1(X11)|~v1_funct_1(X11)|(~r3_wellord1(X9,X10,X11)|epred1_2(X10,X9))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).
fof(c_0_25, negated_conjecture, (v1_relat_1(esk1_0)&(v1_relat_1(esk2_0)&((v1_relat_1(esk3_0)&v1_funct_1(esk3_0))&((v2_wellord1(esk1_0)&r3_wellord1(esk1_0,esk2_0,esk3_0))&~v2_wellord1(esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])).
fof(c_0_26, plain, ![X7]:((~v4_relat_2(X7)|r4_relat_2(X7,k3_relat_1(X7))|~v1_relat_1(X7))&(~r4_relat_2(X7,k3_relat_1(X7))|v4_relat_2(X7)|~v1_relat_1(X7))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_relat_2])])])).
cnf(c_0_27, plain, (r2_wellord1(X1,k3_relat_1(X1))|~v1_wellord1(X1)|~v6_relat_2(X1)|~v8_relat_2(X1)|~r1_relat_2(X1,k3_relat_1(X1))|~r4_relat_2(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_28, plain, (r1_relat_2(X1,k3_relat_1(X1))|~v1_relat_2(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_29, plain, ![X18, X19]:((~v1_relat_2(X18)|v1_relat_2(X19)|~epred1_2(X19,X18))&((~v8_relat_2(X18)|v8_relat_2(X19)|~epred1_2(X19,X18))&((~v6_relat_2(X18)|v6_relat_2(X19)|~epred1_2(X19,X18))&((~v4_relat_2(X18)|v4_relat_2(X19)|~epred1_2(X19,X18))&(~v1_wellord1(X18)|v1_wellord1(X19)|~epred1_2(X19,X18)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])).
cnf(c_0_30, plain, (epred1_2(X2,X1)|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X3)|~v1_funct_1(X3)|~r3_wellord1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, negated_conjecture, (r3_wellord1(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_36, plain, (v1_wellord1(X1)|~r1_wellord1(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_37, plain, (r1_wellord1(X1,X2)|~r2_wellord1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_38, plain, ![X17]:((~r2_wellord1(X17,k3_relat_1(X17))|v2_wellord1(X17)|~v1_relat_1(X17))&(~v2_wellord1(X17)|r2_wellord1(X17,k3_relat_1(X17))|~v1_relat_1(X17))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_wellord1])])])).
cnf(c_0_39, plain, (v6_relat_2(X1)|~r6_relat_2(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_40, plain, (r6_relat_2(X1,X2)|~r2_wellord1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_41, plain, (v8_relat_2(X1)|~r8_relat_2(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_42, plain, (r8_relat_2(X1,X2)|~r2_wellord1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_43, plain, (v1_relat_2(X1)|~r1_relat_2(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_44, plain, (r1_relat_2(X1,X2)|~r2_wellord1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_45, plain, (v4_relat_2(X1)|~r4_relat_2(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_46, plain, (r4_relat_2(X1,X2)|~r2_wellord1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_47, plain, (r2_wellord1(X1,k3_relat_1(X1))|~v1_wellord1(X1)|~v6_relat_2(X1)|~v8_relat_2(X1)|~v1_relat_2(X1)|~r4_relat_2(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_48, plain, (r4_relat_2(X1,k3_relat_1(X1))|~v4_relat_2(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_49, plain, (v1_wellord1(X2)|~v1_wellord1(X1)|~epred1_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_50, negated_conjecture, (epred1_2(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35])])).
cnf(c_0_51, plain, (v1_wellord1(X1)|~r2_wellord1(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_52, plain, (r2_wellord1(X1,k3_relat_1(X1))|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_53, plain, (v6_relat_2(X2)|~v6_relat_2(X1)|~epred1_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_54, plain, (v6_relat_2(X1)|~r2_wellord1(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_55, plain, (v8_relat_2(X2)|~v8_relat_2(X1)|~epred1_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_56, plain, (v8_relat_2(X1)|~r2_wellord1(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_57, plain, (v1_relat_2(X2)|~v1_relat_2(X1)|~epred1_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_58, plain, (v1_relat_2(X1)|~r2_wellord1(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_59, plain, (v4_relat_2(X2)|~v4_relat_2(X1)|~epred1_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_60, plain, (v4_relat_2(X1)|~r2_wellord1(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_61, plain, (v2_wellord1(X1)|~r2_wellord1(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_62, plain, (r2_wellord1(X1,k3_relat_1(X1))|~v1_wellord1(X1)|~v6_relat_2(X1)|~v8_relat_2(X1)|~v1_relat_2(X1)|~v4_relat_2(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_63, plain, (v1_wellord1(esk2_0)|~v1_wellord1(esk1_0)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_64, plain, (v1_wellord1(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_65, negated_conjecture, (v2_wellord1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_66, plain, (v6_relat_2(esk2_0)|~v6_relat_2(esk1_0)), inference(spm,[status(thm)],[c_0_53, c_0_50])).
cnf(c_0_67, plain, (v6_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_54, c_0_52])).
cnf(c_0_68, plain, (v8_relat_2(esk2_0)|~v8_relat_2(esk1_0)), inference(spm,[status(thm)],[c_0_55, c_0_50])).
cnf(c_0_69, plain, (v8_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_56, c_0_52])).
cnf(c_0_70, plain, (v1_relat_2(esk2_0)|~v1_relat_2(esk1_0)), inference(spm,[status(thm)],[c_0_57, c_0_50])).
cnf(c_0_71, plain, (v1_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_58, c_0_52])).
cnf(c_0_72, plain, (v4_relat_2(esk2_0)|~v4_relat_2(esk1_0)), inference(spm,[status(thm)],[c_0_59, c_0_50])).
cnf(c_0_73, plain, (v4_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_60, c_0_52])).
cnf(c_0_74, plain, (v2_wellord1(X1)|~v1_wellord1(X1)|~v6_relat_2(X1)|~v8_relat_2(X1)|~v1_relat_2(X1)|~v4_relat_2(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_75, plain, (v1_wellord1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_65]), c_0_35])])).
cnf(c_0_76, plain, (v6_relat_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_65]), c_0_35])])).
cnf(c_0_77, plain, (v8_relat_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_65]), c_0_35])])).
cnf(c_0_78, plain, (v1_relat_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_65]), c_0_35])])).
cnf(c_0_79, plain, (v4_relat_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_65]), c_0_35])])).
cnf(c_0_80, negated_conjecture, (~v2_wellord1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_81, plain, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79]), c_0_34])]), c_0_80]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 82
# Proof object clause steps            : 59
# Proof object formula steps           : 23
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 28
# Proof object simplifying inferences  : 27
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 31
# Processed clauses                    : 89
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 89
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 5
# Generated clauses                    : 45
# ...of the previous two non-trivial   : 27
# Contextual simplify-reflections      : 0
# Paramodulations                      : 45
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 53
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 40
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 36
# Clause-clause subsumption calls (NU) : 387
# Rec. Clause-clause subsumption calls : 158
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2794

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.003 s
# Total time               : 0.011 s
# Maximum resident set size: 28616 pages
