# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t134_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4)))), file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', t134_zfmisc_1)).
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', d2_zfmisc_1)).
fof(t120_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&k2_zfmisc_1(X3,k2_xboole_0(X1,X2))=k2_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))), file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', t120_zfmisc_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', commutativity_k2_xboole_0)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', t4_xboole_0)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', d10_xboole_0)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', idempotence_k3_xboole_0)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', t12_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', t28_xboole_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', t7_xboole_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', t113_zfmisc_1)).
fof(t123_zfmisc_1, axiom, ![X1, X2, X3, X4]:k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)), file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', t123_zfmisc_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', commutativity_k3_xboole_0)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', d7_xboole_0)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', t7_xboole_0)).
fof(t117_zfmisc_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&((r1_tarski(k2_zfmisc_1(X2,X1),k2_zfmisc_1(X3,X1))|r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))&~(r1_tarski(X2,X3))))), file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', t117_zfmisc_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', l32_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', t100_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', t48_xboole_1)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', t2_boole)).
fof(t5_boole, axiom, ![X1]:k5_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/157183def70b4821b4db51a7ef8508cf.E_input', t5_boole)).
fof(c_0_21, negated_conjecture, ~(![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4))))), inference(assume_negation,[status(cth)],[t134_zfmisc_1])).
fof(c_0_22, plain, ![X7, X8, X9, X10, X13, X14, X15, X16, X17, X18, X20, X21]:((((r2_hidden(esk1_4(X7,X8,X9,X10),X7)|~r2_hidden(X10,X9)|X9!=k2_zfmisc_1(X7,X8))&((r2_hidden(esk2_4(X7,X8,X9,X10),X8)|~r2_hidden(X10,X9)|X9!=k2_zfmisc_1(X7,X8))&(X10=k4_tarski(esk1_4(X7,X8,X9,X10),esk2_4(X7,X8,X9,X10))|~r2_hidden(X10,X9)|X9!=k2_zfmisc_1(X7,X8))))&(~r2_hidden(X14,X7)|(~r2_hidden(X15,X8)|X13!=k4_tarski(X14,X15))|r2_hidden(X13,X9)|X9!=k2_zfmisc_1(X7,X8)))&((~r2_hidden(esk3_3(X16,X17,X18),X18)|(~r2_hidden(X20,X16)|(~r2_hidden(X21,X17)|esk3_3(X16,X17,X18)!=k4_tarski(X20,X21)))|X18=k2_zfmisc_1(X16,X17))&((r2_hidden(esk4_3(X16,X17,X18),X16)|r2_hidden(esk3_3(X16,X17,X18),X18)|X18=k2_zfmisc_1(X16,X17))&((r2_hidden(esk5_3(X16,X17,X18),X17)|r2_hidden(esk3_3(X16,X17,X18),X18)|X18=k2_zfmisc_1(X16,X17))&(esk3_3(X16,X17,X18)=k4_tarski(esk4_3(X16,X17,X18),esk5_3(X16,X17,X18))|r2_hidden(esk3_3(X16,X17,X18),X18)|X18=k2_zfmisc_1(X16,X17)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
fof(c_0_23, plain, ![X35, X36, X37]:(k2_zfmisc_1(k2_xboole_0(X35,X36),X37)=k2_xboole_0(k2_zfmisc_1(X35,X37),k2_zfmisc_1(X36,X37))&k2_zfmisc_1(X37,k2_xboole_0(X35,X36))=k2_xboole_0(k2_zfmisc_1(X37,X35),k2_zfmisc_1(X37,X36))), inference(variable_rename,[status(thm)],[t120_zfmisc_1])).
fof(c_0_24, negated_conjecture, (k2_zfmisc_1(esk7_0,esk8_0)=k2_zfmisc_1(esk9_0,esk10_0)&(esk7_0!=k1_xboole_0&(esk8_0!=k1_xboole_0&(esk7_0!=esk9_0|esk8_0!=esk10_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).
fof(c_0_25, plain, ![X40, X41]:k2_xboole_0(X40,X41)=k2_xboole_0(X41,X40), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_26, plain, ![X24, X25, X27, X28, X29]:((r1_xboole_0(X24,X25)|r2_hidden(esk6_2(X24,X25),k3_xboole_0(X24,X25)))&(~r2_hidden(X29,k3_xboole_0(X27,X28))|~r1_xboole_0(X27,X28))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
cnf(c_0_27, plain, (r2_hidden(esk2_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, plain, (k2_zfmisc_1(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, negated_conjecture, (k2_zfmisc_1(esk7_0,esk8_0)=k2_zfmisc_1(esk9_0,esk10_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_31, plain, ![X61, X62]:(((r1_tarski(X61,X62)|X61!=X62)&(r1_tarski(X62,X61)|X61!=X62))&(~r1_tarski(X61,X62)|~r1_tarski(X62,X61)|X61=X62)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_32, plain, (r2_hidden(esk1_4(X1,X2,X3,X4),X1)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (r2_hidden(esk2_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X2)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_27])).
fof(c_0_35, plain, ![X34]:k3_xboole_0(X34,X34)=X34, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_36, negated_conjecture, (k2_xboole_0(k2_zfmisc_1(esk7_0,esk8_0),k2_zfmisc_1(esk9_0,X1))=k2_zfmisc_1(esk9_0,k2_xboole_0(X1,esk10_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])).
fof(c_0_37, plain, ![X50, X51]:(~r1_tarski(X50,X51)|k2_xboole_0(X50,X51)=X51), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_38, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_39, plain, ![X52, X53]:(~r1_tarski(X52,X53)|k3_xboole_0(X52,X53)=X52), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_40, plain, ![X38, X39]:r1_tarski(X38,k2_xboole_0(X38,X39)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_41, plain, (r2_hidden(esk1_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X1)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_42, plain, (~r1_xboole_0(X1,X2)|~r2_hidden(X3,k2_zfmisc_1(X4,k3_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_43, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_44, plain, ![X56, X57]:((k2_zfmisc_1(X56,X57)!=k1_xboole_0|(X56=k1_xboole_0|X57=k1_xboole_0))&((X56!=k1_xboole_0|k2_zfmisc_1(X56,X57)=k1_xboole_0)&(X57!=k1_xboole_0|k2_zfmisc_1(X56,X57)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_45, negated_conjecture, (k2_zfmisc_1(esk9_0,k2_xboole_0(esk10_0,esk10_0))=k2_zfmisc_1(esk7_0,k2_xboole_0(esk8_0,esk8_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_29]), c_0_28])).
cnf(c_0_46, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_47, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_38])).
fof(c_0_48, plain, ![X42, X43, X44, X45]:k2_zfmisc_1(k3_xboole_0(X42,X43),k3_xboole_0(X44,X45))=k3_xboole_0(k2_zfmisc_1(X42,X44),k2_zfmisc_1(X43,X45)), inference(variable_rename,[status(thm)],[t123_zfmisc_1])).
fof(c_0_49, plain, ![X46, X47]:k3_xboole_0(X46,X47)=k3_xboole_0(X47,X46), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_50, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_52, plain, (~r1_xboole_0(X1,X2)|~r2_hidden(X3,k2_zfmisc_1(k3_xboole_0(X1,X2),X4))), inference(spm,[status(thm)],[c_0_33, c_0_41])).
cnf(c_0_53, plain, (~r1_xboole_0(X1,X1)|~r2_hidden(X2,k2_zfmisc_1(X3,X1))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
fof(c_0_54, plain, ![X48, X49]:((~r1_xboole_0(X48,X49)|k3_xboole_0(X48,X49)=k1_xboole_0)&(k3_xboole_0(X48,X49)!=k1_xboole_0|r1_xboole_0(X48,X49))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_55, plain, (X1=k1_xboole_0|X2=k1_xboole_0|k2_zfmisc_1(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_56, negated_conjecture, (k2_zfmisc_1(esk7_0,k2_xboole_0(esk8_0,esk8_0))=k2_zfmisc_1(esk7_0,esk8_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_29]), c_0_47])])).
cnf(c_0_57, negated_conjecture, (esk7_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_58, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_60, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_61, plain, (k2_zfmisc_1(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_62, plain, (~r1_xboole_0(X1,X1)|~r2_hidden(X2,k2_zfmisc_1(X1,X3))), inference(spm,[status(thm)],[c_0_52, c_0_43])).
cnf(c_0_63, negated_conjecture, (~r1_xboole_0(esk10_0,esk10_0)|~r2_hidden(X1,k2_zfmisc_1(esk7_0,esk8_0))), inference(spm,[status(thm)],[c_0_53, c_0_29])).
cnf(c_0_64, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_54])).
fof(c_0_65, plain, ![X54]:(X54=k1_xboole_0|r2_hidden(esk11_1(X54),X54)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_66, negated_conjecture, (k2_xboole_0(esk8_0,esk8_0)=k1_xboole_0|k2_zfmisc_1(esk7_0,esk8_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57])).
cnf(c_0_67, negated_conjecture, (esk8_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_68, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))=k2_zfmisc_1(k3_xboole_0(X1,X3),X2)), inference(spm,[status(thm)],[c_0_58, c_0_43])).
cnf(c_0_69, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X4),k2_zfmisc_1(X3,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_58])).
cnf(c_0_70, plain, (k3_xboole_0(X1,k2_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_60, c_0_30])).
cnf(c_0_71, negated_conjecture, (k2_xboole_0(k2_zfmisc_1(X1,esk10_0),k2_zfmisc_1(esk7_0,esk8_0))=k2_zfmisc_1(k2_xboole_0(X1,esk9_0),esk10_0)), inference(spm,[status(thm)],[c_0_61, c_0_29])).
cnf(c_0_72, negated_conjecture, (~r1_xboole_0(esk9_0,esk9_0)|~r2_hidden(X1,k2_zfmisc_1(esk7_0,esk8_0))), inference(spm,[status(thm)],[c_0_62, c_0_29])).
fof(c_0_73, plain, ![X58, X59, X60]:((~r1_tarski(k2_zfmisc_1(X59,X58),k2_zfmisc_1(X60,X58))|r1_tarski(X59,X60)|X58=k1_xboole_0)&(~r1_tarski(k2_zfmisc_1(X58,X59),k2_zfmisc_1(X58,X60))|r1_tarski(X59,X60)|X58=k1_xboole_0)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t117_zfmisc_1])])])])).
cnf(c_0_74, negated_conjecture, (esk10_0!=k1_xboole_0|~r2_hidden(X1,k2_zfmisc_1(esk7_0,esk8_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_43])).
cnf(c_0_75, plain, (X1=k1_xboole_0|r2_hidden(esk11_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_76, negated_conjecture, (k2_zfmisc_1(esk7_0,esk8_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_66]), c_0_47])]), c_0_67])).
cnf(c_0_77, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))=k2_zfmisc_1(X1,k3_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_58, c_0_43])).
cnf(c_0_78, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(X1,esk8_0),k2_zfmisc_1(esk7_0,esk10_0))=k2_zfmisc_1(k3_xboole_0(X1,esk9_0),esk10_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_29]), c_0_69])).
cnf(c_0_79, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(k2_xboole_0(X1,X3),X4))=k2_zfmisc_1(X1,k3_xboole_0(X2,X4))), inference(spm,[status(thm)],[c_0_58, c_0_60])).
cnf(c_0_80, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk7_0,esk8_0),k2_zfmisc_1(k2_xboole_0(X1,esk9_0),esk10_0))=k2_zfmisc_1(esk7_0,esk8_0)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_81, negated_conjecture, (esk9_0!=k1_xboole_0|~r2_hidden(X1,k2_zfmisc_1(esk7_0,esk8_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_64]), c_0_43])).
cnf(c_0_82, plain, (r1_tarski(X1,X3)|X2=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_83, negated_conjecture, (esk10_0!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76])).
cnf(c_0_84, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk7_0,esk9_0),esk10_0)=k2_zfmisc_1(esk7_0,k3_xboole_0(esk8_0,esk10_0))), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_85, negated_conjecture, (k2_zfmisc_1(esk7_0,k3_xboole_0(esk8_0,esk10_0))=k2_zfmisc_1(esk7_0,esk8_0)), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_86, plain, (r1_tarski(X2,X3)|X1=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_87, negated_conjecture, (esk9_0!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_75]), c_0_76])).
cnf(c_0_88, negated_conjecture, (k2_zfmisc_1(k2_xboole_0(esk7_0,esk9_0),esk8_0)=k2_zfmisc_1(esk9_0,k2_xboole_0(esk8_0,esk10_0))), inference(spm,[status(thm)],[c_0_61, c_0_36])).
cnf(c_0_89, plain, (k2_xboole_0(X1,X2)=X1|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_46, c_0_30])).
fof(c_0_90, plain, ![X63, X64]:((k4_xboole_0(X63,X64)!=k1_xboole_0|r1_tarski(X63,X64))&(~r1_tarski(X63,X64)|k4_xboole_0(X63,X64)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
fof(c_0_91, plain, ![X65, X66]:k4_xboole_0(X65,X66)=k5_xboole_0(X65,k3_xboole_0(X65,X66)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_92, plain, ![X67, X68]:k4_xboole_0(X67,k4_xboole_0(X67,X68))=k3_xboole_0(X67,X68), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_93, negated_conjecture, (r1_tarski(esk9_0,X1)|~r1_tarski(k2_zfmisc_1(esk7_0,esk8_0),k2_zfmisc_1(X1,esk10_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_29]), c_0_83])).
cnf(c_0_94, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk7_0,esk9_0),esk10_0)=k2_zfmisc_1(esk7_0,esk8_0)), inference(rw,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_95, negated_conjecture, (r1_tarski(X1,esk9_0)|~r1_tarski(k2_zfmisc_1(X1,esk10_0),k2_zfmisc_1(esk7_0,esk8_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_29]), c_0_83])).
cnf(c_0_96, negated_conjecture, (r1_tarski(X1,esk10_0)|~r1_tarski(k2_zfmisc_1(esk9_0,X1),k2_zfmisc_1(esk7_0,esk8_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_29]), c_0_87])).
cnf(c_0_97, negated_conjecture, (k2_zfmisc_1(esk9_0,k2_xboole_0(esk8_0,esk10_0))=k2_zfmisc_1(esk7_0,esk8_0)|~r1_tarski(esk9_0,esk7_0)), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_98, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_99, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_100, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_92])).
fof(c_0_101, plain, ![X69]:k3_xboole_0(X69,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_102, plain, ![X70]:k5_xboole_0(X70,k1_xboole_0)=X70, inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_103, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk7_0,esk8_0),k2_zfmisc_1(esk9_0,k2_xboole_0(X1,esk10_0)))=k2_zfmisc_1(esk7_0,esk8_0)), inference(spm,[status(thm)],[c_0_60, c_0_36])).
cnf(c_0_104, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,k2_xboole_0(X2,X4)))=k2_zfmisc_1(k3_xboole_0(X1,X3),X2)), inference(spm,[status(thm)],[c_0_58, c_0_60])).
cnf(c_0_105, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_106, negated_conjecture, (r1_tarski(esk9_0,k3_xboole_0(esk7_0,esk9_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_47])])).
cnf(c_0_107, negated_conjecture, (r1_tarski(k3_xboole_0(esk7_0,esk9_0),esk9_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_94]), c_0_47])])).
cnf(c_0_108, negated_conjecture, (r1_tarski(k2_xboole_0(esk8_0,esk10_0),esk10_0)|~r1_tarski(esk9_0,esk7_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_47])])).
cnf(c_0_109, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_51, c_0_30])).
cnf(c_0_110, plain, (r1_tarski(X1,X2)|k5_xboole_0(X1,k3_xboole_0(X1,X2))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_111, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2))))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_100, c_0_99]), c_0_99])).
cnf(c_0_112, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_113, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_114, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(k2_xboole_0(X1,X3),X2))), inference(spm,[status(thm)],[c_0_51, c_0_61])).
cnf(c_0_115, negated_conjecture, (k2_zfmisc_1(k2_xboole_0(esk7_0,esk9_0),esk10_0)=k2_zfmisc_1(esk7_0,k2_xboole_0(esk8_0,esk10_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_71]), c_0_30])).
cnf(c_0_116, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk7_0,esk9_0),esk8_0)=k2_zfmisc_1(esk7_0,esk8_0)), inference(spm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_117, negated_conjecture, (k3_xboole_0(esk7_0,esk9_0)=esk9_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_107])])).
cnf(c_0_118, negated_conjecture, (k2_xboole_0(esk8_0,esk10_0)=esk10_0|~r1_tarski(esk9_0,esk7_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_108]), c_0_109])])).
cnf(c_0_119, plain, (r1_tarski(X1,X2)|k5_xboole_0(X1,k3_xboole_0(X2,X1))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_110, c_0_59])).
cnf(c_0_120, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_113]), c_0_43])).
cnf(c_0_121, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk7_0,esk10_0),k2_zfmisc_1(esk7_0,k2_xboole_0(esk8_0,esk10_0)))), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_122, negated_conjecture, (r1_tarski(esk10_0,X1)|~r1_tarski(k2_zfmisc_1(esk7_0,esk8_0),k2_zfmisc_1(esk9_0,X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_29]), c_0_87])).
cnf(c_0_123, negated_conjecture, (k2_zfmisc_1(esk9_0,esk8_0)=k2_zfmisc_1(esk7_0,esk8_0)), inference(rw,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_124, negated_conjecture, (k3_xboole_0(esk8_0,esk10_0)=esk8_0|~r1_tarski(esk9_0,esk7_0)), inference(spm,[status(thm)],[c_0_60, c_0_118])).
cnf(c_0_125, negated_conjecture, (r1_tarski(esk9_0,esk7_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_117]), c_0_120])])).
cnf(c_0_126, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk7_0,esk10_0),k2_zfmisc_1(esk7_0,esk8_0))|~r1_tarski(esk10_0,esk8_0)), inference(spm,[status(thm)],[c_0_121, c_0_89])).
cnf(c_0_127, negated_conjecture, (r1_tarski(esk10_0,esk8_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_47])])).
cnf(c_0_128, negated_conjecture, (k3_xboole_0(esk8_0,esk10_0)=esk8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_124, c_0_125])])).
cnf(c_0_129, negated_conjecture, (r1_tarski(esk7_0,esk9_0)|~r1_tarski(esk10_0,esk8_0)), inference(spm,[status(thm)],[c_0_95, c_0_126])).
cnf(c_0_130, negated_conjecture, (esk7_0!=esk9_0|esk8_0!=esk10_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_131, negated_conjecture, (esk10_0=esk8_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_127]), c_0_59]), c_0_128])).
cnf(c_0_132, negated_conjecture, (r1_tarski(esk7_0,esk9_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_129, c_0_127])])).
cnf(c_0_133, negated_conjecture, (esk9_0!=esk7_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_130, c_0_131])])).
cnf(c_0_134, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_132]), c_0_117]), c_0_133]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 135
# Proof object clause steps            : 92
# Proof object formula steps           : 43
# Proof object conjectures             : 50
# Proof object clause conjectures      : 47
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 21
# Proof object generating inferences   : 54
# Proof object simplifying inferences  : 54
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 39
# Processed clauses                    : 4302
# ...of these trivial                  : 63
# ...subsumed                          : 3393
# ...remaining for further processing  : 846
# Other redundant clauses eliminated   : 126
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 44
# Backward-rewritten                   : 432
# Generated clauses                    : 29357
# ...of the previous two non-trivial   : 27187
# Contextual simplify-reflections      : 1
# Paramodulations                      : 29229
# Factorizations                       : 3
# Equation resolutions                 : 126
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 318
#    Positive orientable unit clauses  : 53
#    Positive unorientable unit clauses: 4
#    Negative unit clauses             : 15
#    Non-unit-clauses                  : 246
# Current number of unprocessed clauses: 21952
# ...number of literals in the above   : 57530
# Current number of archived formulas  : 0
# Current number of archived clauses   : 515
# Clause-clause subsumption calls (NU) : 63958
# Rec. Clause-clause subsumption calls : 38505
# Non-unit clause-clause subsumptions  : 1423
# Unit Clause-clause subsumption calls : 2284
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 264
# BW rewrite match successes           : 133
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 397067

# -------------------------------------------------
# User time                : 0.244 s
# System time              : 0.007 s
# Total time               : 0.250 s
# Maximum resident set size: 28636 pages
