# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t22_ordinal1, axiom, ![X1]:(v1_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>![X3]:(v3_ordinal1(X3)=>((r1_tarski(X1,X2)&r2_hidden(X2,X3))=>r2_hidden(X1,X3))))), file('data/MPTP/proofs/1584ab4880fa46a3926225d36de6ac0c.E_input', t22_ordinal1)).
fof(t23_ordinal1, axiom, ![X1, X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)=>v3_ordinal1(X1))), file('data/MPTP/proofs/1584ab4880fa46a3926225d36de6ac0c.E_input', t23_ordinal1)).
fof(redefinition_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)<=>r1_tarski(X1,X2))), file('data/MPTP/proofs/1584ab4880fa46a3926225d36de6ac0c.E_input', redefinition_r1_ordinal1)).
fof(t26_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r1_ordinal1(X1,X2)|r2_hidden(X2,X1)))), file('data/MPTP/proofs/1584ab4880fa46a3926225d36de6ac0c.E_input', t26_ordinal1)).
fof(t13_ordinal1, axiom, ![X1, X2]:(r2_hidden(X1,k1_ordinal1(X2))<=>(r2_hidden(X1,X2)|X1=X2)), file('data/MPTP/proofs/1584ab4880fa46a3926225d36de6ac0c.E_input', t13_ordinal1)).
fof(d1_ordinal1, axiom, ![X1]:k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1)), file('data/MPTP/proofs/1584ab4880fa46a3926225d36de6ac0c.E_input', d1_ordinal1)).
fof(cc1_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>(v1_ordinal1(X1)&v2_ordinal1(X1))), file('data/MPTP/proofs/1584ab4880fa46a3926225d36de6ac0c.E_input', cc1_ordinal1)).
fof(t29_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>v3_ordinal1(k1_ordinal1(X1))), file('data/MPTP/proofs/1584ab4880fa46a3926225d36de6ac0c.E_input', t29_ordinal1)).
fof(t31_ordinal1, conjecture, ![X1]:(![X2]:(r2_hidden(X2,X1)=>(v3_ordinal1(X2)&r1_tarski(X2,X1)))=>v3_ordinal1(X1)), file('data/MPTP/proofs/1584ab4880fa46a3926225d36de6ac0c.E_input', t31_ordinal1)).
fof(d2_ordinal1, axiom, ![X1]:(v1_ordinal1(X1)<=>![X2]:(r2_hidden(X2,X1)=>r1_tarski(X2,X1))), file('data/MPTP/proofs/1584ab4880fa46a3926225d36de6ac0c.E_input', d2_ordinal1)).
fof(cc2_ordinal1, axiom, ![X1]:((v1_ordinal1(X1)&v2_ordinal1(X1))=>v3_ordinal1(X1)), file('data/MPTP/proofs/1584ab4880fa46a3926225d36de6ac0c.E_input', cc2_ordinal1)).
fof(d3_ordinal1, axiom, ![X1]:(v2_ordinal1(X1)<=>![X2, X3]:~((r2_hidden(X2,X1)&(r2_hidden(X3,X1)&(~(r2_hidden(X2,X3))&(~(X2=X3)&~(r2_hidden(X3,X2)))))))), file('data/MPTP/proofs/1584ab4880fa46a3926225d36de6ac0c.E_input', d3_ordinal1)).
fof(c_0_12, plain, ![X4, X5, X6]:(~v1_ordinal1(X4)|(~v3_ordinal1(X5)|(~v3_ordinal1(X6)|(~r1_tarski(X4,X5)|~r2_hidden(X5,X6)|r2_hidden(X4,X6))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t22_ordinal1])])])).
fof(c_0_13, plain, ![X7, X8]:(~v3_ordinal1(X8)|(~r2_hidden(X7,X8)|v3_ordinal1(X7))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_ordinal1])])).
fof(c_0_14, plain, ![X9, X10]:((~r1_ordinal1(X9,X10)|r1_tarski(X9,X10)|(~v3_ordinal1(X9)|~v3_ordinal1(X10)))&(~r1_tarski(X9,X10)|r1_ordinal1(X9,X10)|(~v3_ordinal1(X9)|~v3_ordinal1(X10)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r1_ordinal1])])])).
fof(c_0_15, plain, ![X11, X12]:(~v3_ordinal1(X11)|(~v3_ordinal1(X12)|(r1_ordinal1(X11,X12)|r2_hidden(X12,X11)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t26_ordinal1])])])).
fof(c_0_16, plain, ![X13, X14]:((~r2_hidden(X13,k1_ordinal1(X14))|(r2_hidden(X13,X14)|X13=X14))&((~r2_hidden(X13,X14)|r2_hidden(X13,k1_ordinal1(X14)))&(X13!=X14|r2_hidden(X13,k1_ordinal1(X14))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_ordinal1])])])).
fof(c_0_17, plain, ![X15]:k1_ordinal1(X15)=k2_xboole_0(X15,k1_tarski(X15)), inference(variable_rename,[status(thm)],[d1_ordinal1])).
cnf(c_0_18, plain, (r2_hidden(X1,X3)|~v1_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~r1_tarski(X1,X2)|~r2_hidden(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, plain, (r1_tarski(X1,X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (r1_ordinal1(X1,X2)|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_22, plain, ![X16]:((v1_ordinal1(X16)|~v3_ordinal1(X16))&(v2_ordinal1(X16)|~v3_ordinal1(X16))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_ordinal1])])])).
cnf(c_0_23, plain, (r2_hidden(X1,k1_ordinal1(X2))|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (r2_hidden(X1,X2)|~r2_hidden(X3,X2)|~r1_tarski(X1,X3)|~v3_ordinal1(X2)|~v1_ordinal1(X1)), inference(csr,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_26, plain, (r2_hidden(X1,X2)|r1_tarski(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27, plain, (v1_ordinal1(X1)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, plain, (r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))|X1!=X2), inference(rw,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_29, plain, ![X17]:(~v3_ordinal1(X17)|v3_ordinal1(k1_ordinal1(X17))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t29_ordinal1])])).
fof(c_0_30, negated_conjecture, ~(![X1]:(![X2]:(r2_hidden(X2,X1)=>(v3_ordinal1(X2)&r1_tarski(X2,X1)))=>v3_ordinal1(X1))), inference(assume_negation,[status(cth)],[t31_ordinal1])).
cnf(c_0_31, plain, (r2_hidden(X1,X2)|r2_hidden(X2,X3)|~r2_hidden(X1,X3)|~v3_ordinal1(X3)|~v3_ordinal1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_19])).
cnf(c_0_32, plain, (r2_hidden(X1,k2_xboole_0(X1,k1_tarski(X1)))), inference(er,[status(thm)],[c_0_28])).
cnf(c_0_33, plain, (v3_ordinal1(k1_ordinal1(X1))|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_34, plain, ![X20, X21, X22]:((~v1_ordinal1(X20)|(~r2_hidden(X21,X20)|r1_tarski(X21,X20)))&((r2_hidden(esk2_1(X22),X22)|v1_ordinal1(X22))&(~r1_tarski(esk2_1(X22),X22)|v1_ordinal1(X22)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_ordinal1])])])])])])).
fof(c_0_35, negated_conjecture, ![X19]:(((v3_ordinal1(X19)|~r2_hidden(X19,esk1_0))&(r1_tarski(X19,esk1_0)|~r2_hidden(X19,esk1_0)))&~v3_ordinal1(esk1_0)), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])])])).
cnf(c_0_36, plain, (r2_hidden(X1,X2)|X1=X2|~r2_hidden(X1,k1_ordinal1(X2))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_37, plain, (r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))|r2_hidden(X2,X1)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_38, plain, (v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)), inference(rw,[status(thm)],[c_0_33, c_0_24])).
cnf(c_0_39, plain, (v1_ordinal1(X1)|~r1_tarski(esk2_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_40, negated_conjecture, (r1_tarski(X1,esk1_0)|~r2_hidden(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41, plain, (X1=X2|r2_hidden(X1,X2)|~r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))), inference(rw,[status(thm)],[c_0_36, c_0_24])).
cnf(c_0_42, plain, (r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
fof(c_0_43, plain, ![X24]:(~v1_ordinal1(X24)|~v2_ordinal1(X24)|v3_ordinal1(X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_ordinal1])])).
cnf(c_0_44, negated_conjecture, (v1_ordinal1(esk1_0)|~r2_hidden(esk2_1(esk1_0),esk1_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_45, plain, (r2_hidden(esk2_1(X1),X1)|v1_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_46, plain, (X1=X2|r2_hidden(X2,X1)|r2_hidden(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_47, negated_conjecture, (v3_ordinal1(X1)|~r2_hidden(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_48, plain, ![X25, X26, X27, X28]:((~v2_ordinal1(X25)|(~r2_hidden(X26,X25)|(~r2_hidden(X27,X25)|(r2_hidden(X26,X27)|(X26=X27|r2_hidden(X27,X26))))))&((r2_hidden(esk3_1(X28),X28)|v2_ordinal1(X28))&((r2_hidden(esk4_1(X28),X28)|v2_ordinal1(X28))&((~r2_hidden(esk3_1(X28),esk4_1(X28))|v2_ordinal1(X28))&((esk3_1(X28)!=esk4_1(X28)|v2_ordinal1(X28))&(~r2_hidden(esk4_1(X28),esk3_1(X28))|v2_ordinal1(X28))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d3_ordinal1])])])])])])])).
cnf(c_0_49, plain, (v3_ordinal1(X1)|~v1_ordinal1(X1)|~v2_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_50, negated_conjecture, (v1_ordinal1(esk1_0)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_51, negated_conjecture, (~v3_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_52, negated_conjecture, (X1=X2|r2_hidden(X1,X2)|r2_hidden(X2,X1)|~r2_hidden(X1,esk1_0)|~v3_ordinal1(X2)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_53, plain, (r2_hidden(esk4_1(X1),X1)|v2_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_54, negated_conjecture, (~v2_ordinal1(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51])).
cnf(c_0_55, negated_conjecture, (esk4_1(esk1_0)=X1|r2_hidden(X1,esk4_1(esk1_0))|r2_hidden(esk4_1(esk1_0),X1)|~v3_ordinal1(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54])).
cnf(c_0_56, negated_conjecture, (esk4_1(esk1_0)=X1|r2_hidden(esk4_1(esk1_0),X1)|r2_hidden(X1,esk4_1(esk1_0))|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_55, c_0_47])).
cnf(c_0_57, plain, (r2_hidden(esk3_1(X1),X1)|v2_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_58, plain, (v2_ordinal1(X1)|~r2_hidden(esk4_1(X1),esk3_1(X1))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)|r2_hidden(esk3_1(esk1_0),esk4_1(esk1_0))|r2_hidden(esk4_1(esk1_0),esk3_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_54])).
cnf(c_0_60, plain, (v2_ordinal1(X1)|~r2_hidden(esk3_1(X1),esk4_1(X1))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_61, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)|r2_hidden(esk3_1(esk1_0),esk4_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_54])).
cnf(c_0_62, plain, (v2_ordinal1(X1)|esk3_1(X1)!=esk4_1(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_63, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_54])).
cnf(c_0_64, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_54]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 65
# Proof object clause steps            : 40
# Proof object formula steps           : 25
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 13
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 24
# Processed clauses                    : 89
# ...of these trivial                  : 0
# ...subsumed                          : 14
# ...remaining for further processing  : 75
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 4
# Generated clauses                    : 97
# ...of the previous two non-trivial   : 84
# Contextual simplify-reflections      : 3
# Paramodulations                      : 96
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 45
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 39
# Current number of unprocessed clauses: 41
# ...number of literals in the above   : 223
# Current number of archived formulas  : 0
# Current number of archived clauses   : 30
# Clause-clause subsumption calls (NU) : 369
# Rec. Clause-clause subsumption calls : 103
# Non-unit clause-clause subsumptions  : 17
# Unit Clause-clause subsumption calls : 46
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3487

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28592 pages
