# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t19_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k5_waybel_0(X1,X2)=k5_waybel_0(X1,X3)=>X2=X3)))), file('data/MPTP/proofs/1588003384944092a263d0e8961d1b1a.E_input', t19_waybel_0)).
fof(t17_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2))))), file('data/MPTP/proofs/1588003384944092a263d0e8961d1b1a.E_input', t17_waybel_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/1588003384944092a263d0e8961d1b1a.E_input', t24_orders_2)).
fof(t18_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))&((r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))=>(X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3))))))))), file('data/MPTP/proofs/1588003384944092a263d0e8961d1b1a.E_input', t18_yellow_0)).
fof(c_0_4, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k5_waybel_0(X1,X2)=k5_waybel_0(X1,X3)=>X2=X3))))), inference(assume_negation,[status(cth)],[t19_waybel_0])).
fof(c_0_5, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))&((r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))=>(X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3))))))), introduced(definition)).
fof(c_0_6, plain, ![X9, X10, X11]:((~r2_hidden(X11,k5_waybel_0(X9,X10))|r1_orders_2(X9,X11,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(v2_struct_0(X9)|~l1_orders_2(X9)))&(~r1_orders_2(X9,X11,X10)|r2_hidden(X11,k5_waybel_0(X9,X10))|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(v2_struct_0(X9)|~l1_orders_2(X9)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_waybel_0])])])])])).
fof(c_0_7, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(k5_waybel_0(esk1_0,esk2_0)=k5_waybel_0(esk1_0,esk3_0)&esk2_0!=esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_4])])])])).
fof(c_0_8, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))&((r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))=>(X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3))))))), inference(split_equiv,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (r1_orders_2(X2,X1,X3)|v2_struct_0(X2)|~r2_hidden(X1,k5_waybel_0(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, negated_conjecture, (k5_waybel_0(esk1_0,esk2_0)=k5_waybel_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_14, plain, ![X12, X13]:(v2_struct_0(X12)|(~v3_orders_2(X12)|~l1_orders_2(X12))|(~m1_subset_1(X13,u1_struct_0(X12))|r1_orders_2(X12,X13,X13))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_15, plain, ![X17, X18, X19, X20, X21]:(((r1_orders_2(X19,X18,X20)|(X20!=k10_lattice3(X19,X18,X17)|~r1_yellow_0(X19,k2_tarski(X18,X17)))|~m1_subset_1(X20,u1_struct_0(X19))|~epred1_3(X19,X18,X17))&((r1_orders_2(X19,X17,X20)|(X20!=k10_lattice3(X19,X18,X17)|~r1_yellow_0(X19,k2_tarski(X18,X17)))|~m1_subset_1(X20,u1_struct_0(X19))|~epred1_3(X19,X18,X17))&(~m1_subset_1(X21,u1_struct_0(X19))|(~r1_orders_2(X19,X18,X21)|~r1_orders_2(X19,X17,X21)|r1_orders_2(X19,X20,X21))|(X20!=k10_lattice3(X19,X18,X17)|~r1_yellow_0(X19,k2_tarski(X18,X17)))|~m1_subset_1(X20,u1_struct_0(X19))|~epred1_3(X19,X18,X17))))&(((X20=k10_lattice3(X19,X18,X17)|(m1_subset_1(esk4_4(X17,X18,X19,X20),u1_struct_0(X19))|~r1_orders_2(X19,X17,X20)|~r1_orders_2(X19,X18,X20))|~m1_subset_1(X20,u1_struct_0(X19))|~epred1_3(X19,X18,X17))&(r1_yellow_0(X19,k2_tarski(X18,X17))|(m1_subset_1(esk4_4(X17,X18,X19,X20),u1_struct_0(X19))|~r1_orders_2(X19,X17,X20)|~r1_orders_2(X19,X18,X20))|~m1_subset_1(X20,u1_struct_0(X19))|~epred1_3(X19,X18,X17)))&((((X20=k10_lattice3(X19,X18,X17)|(r1_orders_2(X19,X18,esk4_4(X17,X18,X19,X20))|~r1_orders_2(X19,X17,X20)|~r1_orders_2(X19,X18,X20))|~m1_subset_1(X20,u1_struct_0(X19))|~epred1_3(X19,X18,X17))&(r1_yellow_0(X19,k2_tarski(X18,X17))|(r1_orders_2(X19,X18,esk4_4(X17,X18,X19,X20))|~r1_orders_2(X19,X17,X20)|~r1_orders_2(X19,X18,X20))|~m1_subset_1(X20,u1_struct_0(X19))|~epred1_3(X19,X18,X17)))&((X20=k10_lattice3(X19,X18,X17)|(r1_orders_2(X19,X17,esk4_4(X17,X18,X19,X20))|~r1_orders_2(X19,X17,X20)|~r1_orders_2(X19,X18,X20))|~m1_subset_1(X20,u1_struct_0(X19))|~epred1_3(X19,X18,X17))&(r1_yellow_0(X19,k2_tarski(X18,X17))|(r1_orders_2(X19,X17,esk4_4(X17,X18,X19,X20))|~r1_orders_2(X19,X17,X20)|~r1_orders_2(X19,X18,X20))|~m1_subset_1(X20,u1_struct_0(X19))|~epred1_3(X19,X18,X17))))&((X20=k10_lattice3(X19,X18,X17)|(~r1_orders_2(X19,X20,esk4_4(X17,X18,X19,X20))|~r1_orders_2(X19,X17,X20)|~r1_orders_2(X19,X18,X20))|~m1_subset_1(X20,u1_struct_0(X19))|~epred1_3(X19,X18,X17))&(r1_yellow_0(X19,k2_tarski(X18,X17))|(~r1_orders_2(X19,X20,esk4_4(X17,X18,X19,X20))|~r1_orders_2(X19,X17,X20)|~r1_orders_2(X19,X18,X20))|~m1_subset_1(X20,u1_struct_0(X19))|~epred1_3(X19,X18,X17)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])).
cnf(c_0_16, negated_conjecture, (r1_orders_2(esk1_0,X1,esk3_0)|~r2_hidden(X1,k5_waybel_0(esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11]), c_0_12])]), c_0_13])).
cnf(c_0_17, plain, (r2_hidden(X2,k5_waybel_0(X1,X3))|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_18, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_19, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_21, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>epred1_3(X1,X2,X3)))), inference(apply_def,[status(thm)],[t18_yellow_0, c_0_5])).
cnf(c_0_22, plain, (X1=k10_lattice3(X2,X3,X4)|r1_orders_2(X2,X3,esk4_4(X4,X3,X2,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, negated_conjecture, (r1_orders_2(esk1_0,X1,esk3_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18]), c_0_12])]), c_0_13])).
cnf(c_0_24, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_18]), c_0_12]), c_0_20])]), c_0_13])).
fof(c_0_25, plain, ![X14, X15, X16]:(~v5_orders_2(X14)|~l1_orders_2(X14)|(~m1_subset_1(X15,u1_struct_0(X14))|(~m1_subset_1(X16,u1_struct_0(X14))|epred1_3(X14,X15,X16)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).
cnf(c_0_26, negated_conjecture, (esk3_0=k10_lattice3(esk1_0,X1,X2)|r1_orders_2(esk1_0,X1,esk4_4(X2,X1,esk1_0,esk3_0))|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,esk3_0)|~r1_orders_2(esk1_0,X1,esk3_0)), inference(spm,[status(thm)],[c_0_22, c_0_11])).
cnf(c_0_27, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_18]), c_0_24])])).
cnf(c_0_28, plain, (epred1_3(X1,X2,X3)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_29, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_30, plain, (X1=k10_lattice3(X2,X3,X4)|r1_orders_2(X2,X4,esk4_4(X4,X3,X2,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_31, negated_conjecture, (esk3_0=k10_lattice3(esk1_0,X1,esk2_0)|r1_orders_2(esk1_0,X1,esk4_4(esk2_0,X1,esk1_0,esk3_0))|~epred1_3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,X1,esk3_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32, negated_conjecture, (epred1_3(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_12])])).
cnf(c_0_33, negated_conjecture, (r2_hidden(X1,k5_waybel_0(esk1_0,esk2_0))|~r1_orders_2(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_10]), c_0_11]), c_0_12])]), c_0_13])).
cnf(c_0_34, negated_conjecture, (esk2_0=k10_lattice3(esk1_0,X1,X2)|r1_orders_2(esk1_0,X2,esk4_4(X2,X1,esk1_0,esk2_0))|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)), inference(spm,[status(thm)],[c_0_30, c_0_18])).
cnf(c_0_35, negated_conjecture, (esk3_0=k10_lattice3(esk1_0,X1,esk2_0)|r1_orders_2(esk1_0,X1,esk4_4(esk2_0,X1,esk1_0,esk3_0))|~r1_orders_2(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_18])])).
cnf(c_0_36, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_11]), c_0_12]), c_0_20])]), c_0_13])).
cnf(c_0_37, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_33]), c_0_18]), c_0_12])]), c_0_13])).
cnf(c_0_38, negated_conjecture, (k10_lattice3(esk1_0,X1,esk2_0)=esk2_0|r1_orders_2(esk1_0,esk2_0,esk4_4(esk2_0,X1,esk1_0,esk2_0))|~epred1_3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)), inference(spm,[status(thm)],[c_0_34, c_0_24])).
cnf(c_0_39, negated_conjecture, (k10_lattice3(esk1_0,esk3_0,esk2_0)=esk3_0|r1_orders_2(esk1_0,esk3_0,esk4_4(esk2_0,esk3_0,esk1_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_11]), c_0_36])])).
cnf(c_0_40, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_11]), c_0_36])])).
cnf(c_0_41, negated_conjecture, (esk2_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_42, plain, (X1=k10_lattice3(X2,X3,X4)|~r1_orders_2(X2,X1,esk4_4(X4,X3,X2,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_43, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk4_4(esk2_0,esk3_0,esk1_0,esk3_0))|r1_orders_2(esk1_0,esk2_0,esk4_4(esk2_0,esk3_0,esk1_0,esk2_0))|~epred1_3(esk1_0,esk3_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])]), c_0_41])).
cnf(c_0_44, plain, (k10_lattice3(esk1_0,esk3_0,esk2_0)=esk3_0|r1_orders_2(esk1_0,esk2_0,esk4_4(esk2_0,esk3_0,esk1_0,esk2_0))|~epred1_3(esk1_0,esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_27]), c_0_36]), c_0_11])])).
cnf(c_0_45, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk4_4(esk2_0,esk3_0,esk1_0,esk2_0))|~epred1_3(esk1_0,esk3_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_44]), c_0_40])]), c_0_41])).
cnf(c_0_46, plain, (k10_lattice3(esk1_0,esk3_0,esk2_0)=esk2_0|~epred1_3(esk1_0,esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_45]), c_0_24]), c_0_40]), c_0_18])])).
cnf(c_0_47, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk4_4(esk2_0,esk3_0,esk1_0,esk3_0))|~epred1_3(esk1_0,esk3_0,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_46]), c_0_41])).
cnf(c_0_48, plain, (k10_lattice3(esk1_0,esk3_0,esk2_0)=esk3_0|~epred1_3(esk1_0,esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_47]), c_0_27]), c_0_36]), c_0_11])])).
cnf(c_0_49, plain, (~epred1_3(esk1_0,esk3_0,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_48]), c_0_41])).
cnf(c_0_50, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_32]), c_0_18]), c_0_11])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 51
# Proof object clause steps            : 38
# Proof object formula steps           : 13
# Proof object conjectures             : 30
# Proof object clause conjectures      : 27
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 23
# Proof object simplifying inferences  : 57
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 210
# ...of these trivial                  : 7
# ...subsumed                          : 4
# ...remaining for further processing  : 199
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 32
# Backward-rewritten                   : 0
# Generated clauses                    : 795
# ...of the previous two non-trivial   : 678
# Contextual simplify-reflections      : 9
# Paramodulations                      : 795
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 144
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 131
# Current number of unprocessed clauses: 509
# ...number of literals in the above   : 3751
# Current number of archived formulas  : 0
# Current number of archived clauses   : 55
# Clause-clause subsumption calls (NU) : 5093
# Rec. Clause-clause subsumption calls : 544
# Non-unit clause-clause subsumptions  : 45
# Unit Clause-clause subsumption calls : 101
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 30593

# -------------------------------------------------
# User time                : 0.021 s
# System time              : 0.003 s
# Total time               : 0.025 s
# Maximum resident set size: 28800 pages
