# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/15cfd0a56e1d4915adae0b7493eca971.E_input', d5_xboole_0)).
fof(t9_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k4_tarski(X3,X4))))))), file('data/MPTP/proofs/15cfd0a56e1d4915adae0b7493eca971.E_input', t9_mcart_1)).
fof(t20_mcart_1, conjecture, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1)))), file('data/MPTP/proofs/15cfd0a56e1d4915adae0b7493eca971.E_input', t20_mcart_1)).
fof(d1_enumset1, axiom, ![X1, X2, X3, X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3)))))), file('data/MPTP/proofs/15cfd0a56e1d4915adae0b7493eca971.E_input', d1_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/15cfd0a56e1d4915adae0b7493eca971.E_input', t71_enumset1)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/15cfd0a56e1d4915adae0b7493eca971.E_input', t7_mcart_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/15cfd0a56e1d4915adae0b7493eca971.E_input', d1_tarski)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/15cfd0a56e1d4915adae0b7493eca971.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/15cfd0a56e1d4915adae0b7493eca971.E_input', t70_enumset1)).
fof(c_0_9, plain, ![X27, X28, X29, X30, X31, X32, X33, X34]:((((r2_hidden(X30,X27)|~r2_hidden(X30,X29)|X29!=k4_xboole_0(X27,X28))&(~r2_hidden(X30,X28)|~r2_hidden(X30,X29)|X29!=k4_xboole_0(X27,X28)))&(~r2_hidden(X31,X27)|r2_hidden(X31,X28)|r2_hidden(X31,X29)|X29!=k4_xboole_0(X27,X28)))&((~r2_hidden(esk5_3(X32,X33,X34),X34)|(~r2_hidden(esk5_3(X32,X33,X34),X32)|r2_hidden(esk5_3(X32,X33,X34),X33))|X34=k4_xboole_0(X32,X33))&((r2_hidden(esk5_3(X32,X33,X34),X32)|r2_hidden(esk5_3(X32,X33,X34),X34)|X34=k4_xboole_0(X32,X33))&(~r2_hidden(esk5_3(X32,X33,X34),X33)|r2_hidden(esk5_3(X32,X33,X34),X34)|X34=k4_xboole_0(X32,X33))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_10, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_11, plain, ![X50, X52, X53]:((r2_hidden(esk7_1(X50),X50)|X50=k1_xboole_0)&((~r2_hidden(X52,X50)|esk7_1(X50)!=k4_tarski(X52,X53)|X50=k1_xboole_0)&(~r2_hidden(X53,X50)|esk7_1(X50)!=k4_tarski(X52,X53)|X50=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_mcart_1])])])])])])).
cnf(c_0_12, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (r2_hidden(esk7_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_12])).
cnf(c_0_16, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r2_hidden(esk7_1(k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_17, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|r2_hidden(esk7_1(k4_xboole_0(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_15, c_0_14])).
fof(c_0_18, negated_conjecture, ~(![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1))))), inference(assume_negation,[status(cth)],[t20_mcart_1])).
cnf(c_0_19, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_16, c_0_17])).
fof(c_0_20, plain, ![X13, X14, X15, X16, X17, X18, X19, X20, X21, X22]:(((~r2_hidden(X17,X16)|(X17=X13|(X17=X14|X17=X15))|X16!=k1_enumset1(X13,X14,X15))&((X18!=X13|r2_hidden(X18,X16)|X16!=k1_enumset1(X13,X14,X15))&((X18!=X14|r2_hidden(X18,X16)|X16!=k1_enumset1(X13,X14,X15))&(X18!=X15|r2_hidden(X18,X16)|X16!=k1_enumset1(X13,X14,X15)))))&(((esk1_4(X19,X20,X21,X22)!=X19|~r2_hidden(esk1_4(X19,X20,X21,X22),X22)|X22=k1_enumset1(X19,X20,X21))&((esk1_4(X19,X20,X21,X22)!=X20|~r2_hidden(esk1_4(X19,X20,X21,X22),X22)|X22=k1_enumset1(X19,X20,X21))&(esk1_4(X19,X20,X21,X22)!=X21|~r2_hidden(esk1_4(X19,X20,X21,X22),X22)|X22=k1_enumset1(X19,X20,X21))))&(r2_hidden(esk1_4(X19,X20,X21,X22),X22)|(esk1_4(X19,X20,X21,X22)=X19|(esk1_4(X19,X20,X21,X22)=X20|esk1_4(X19,X20,X21,X22)=X21))|X22=k1_enumset1(X19,X20,X21)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])).
fof(c_0_21, plain, ![X10, X11, X12]:k2_enumset1(X10,X10,X11,X12)=k1_enumset1(X10,X11,X12), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_22, plain, ![X38, X39]:(k1_mcart_1(k4_tarski(X38,X39))=X38&k2_mcart_1(k4_tarski(X38,X39))=X39), inference(variable_rename,[status(thm)],[t7_mcart_1])).
fof(c_0_23, negated_conjecture, (esk2_0=k4_tarski(esk3_0,esk4_0)&(esk2_0=k1_mcart_1(esk2_0)|esk2_0=k2_mcart_1(esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_18])])])])).
cnf(c_0_24, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_13, c_0_19])).
cnf(c_0_25, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X5,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, negated_conjecture, (esk2_0=k4_tarski(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X2,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30, plain, (~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_19]), c_0_24])).
cnf(c_0_31, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_enumset1(X4,X4,X5,X2)), inference(rw,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_32, negated_conjecture, (esk3_0=k1_mcart_1(esk2_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_33, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_enumset1(X2,X2,X4,X5)), inference(rw,[status(thm)],[c_0_29, c_0_26])).
cnf(c_0_34, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_30, c_0_17])).
cnf(c_0_35, plain, (r2_hidden(X1,X2)|X2!=k2_enumset1(X3,X3,X4,X1)), inference(er,[status(thm)],[c_0_31])).
cnf(c_0_36, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_37, negated_conjecture, (k4_tarski(k1_mcart_1(esk2_0),esk4_0)=esk2_0), inference(rw,[status(thm)],[c_0_28, c_0_32])).
fof(c_0_38, plain, ![X40, X41, X42, X43, X44, X45]:(((~r2_hidden(X42,X41)|X42=X40|X41!=k1_tarski(X40))&(X43!=X40|r2_hidden(X43,X41)|X41!=k1_tarski(X40)))&((~r2_hidden(esk6_2(X44,X45),X45)|esk6_2(X44,X45)!=X44|X45=k1_tarski(X44))&(r2_hidden(esk6_2(X44,X45),X45)|esk6_2(X44,X45)=X44|X45=k1_tarski(X44)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_39, plain, ![X47]:k2_tarski(X47,X47)=k1_tarski(X47), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_40, plain, ![X48, X49]:k1_enumset1(X48,X48,X49)=k2_tarski(X48,X49), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_41, plain, (r2_hidden(X1,X2)|X2!=k2_enumset1(X1,X1,X3,X4)), inference(er,[status(thm)],[c_0_33])).
cnf(c_0_42, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_34]), c_0_30])).
cnf(c_0_43, plain, (r2_hidden(X1,k2_enumset1(X2,X2,X3,X1))), inference(er,[status(thm)],[c_0_35])).
cnf(c_0_44, negated_conjecture, (esk4_0=k2_mcart_1(esk2_0)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_45, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_46, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_47, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_48, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk7_1(X2)!=k4_tarski(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_49, plain, (r2_hidden(X1,k2_enumset1(X1,X1,X2,X3))), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_50, plain, (k2_enumset1(X1,X1,X2,X3)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_51, negated_conjecture, (k4_tarski(k1_mcart_1(esk2_0),k2_mcart_1(esk2_0))=esk2_0), inference(rw,[status(thm)],[c_0_37, c_0_44])).
cnf(c_0_52, negated_conjecture, (esk2_0=k1_mcart_1(esk2_0)|esk2_0=k2_mcart_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_53, plain, (X1=X3|X2!=k2_enumset1(X3,X3,X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_46]), c_0_47]), c_0_26])).
cnf(c_0_54, plain, (esk7_1(k2_enumset1(X1,X1,X2,X3))!=k4_tarski(X4,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])).
cnf(c_0_55, negated_conjecture, (k4_tarski(k1_mcart_1(esk2_0),esk2_0)=esk2_0|k1_mcart_1(esk2_0)=esk2_0), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_56, plain, (X1=X2|~r2_hidden(X1,k2_enumset1(X2,X2,X2,X2))), inference(er,[status(thm)],[c_0_53])).
cnf(c_0_57, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk7_1(X2)!=k4_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_58, negated_conjecture, (k1_mcart_1(esk2_0)=esk2_0|esk7_1(k2_enumset1(esk2_0,esk2_0,X1,X2))!=esk2_0), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_59, plain, (esk7_1(k2_enumset1(X1,X1,X1,X1))=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_14]), c_0_50])).
cnf(c_0_60, plain, (X1=k1_xboole_0|k4_tarski(esk7_1(X1),X2)!=esk7_1(X1)), inference(spm,[status(thm)],[c_0_57, c_0_14])).
cnf(c_0_61, negated_conjecture, (k1_mcart_1(esk2_0)=esk2_0), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_62, plain, (k4_tarski(X1,X2)!=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_59]), c_0_50])).
cnf(c_0_63, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_61]), c_0_62]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 64
# Proof object clause steps            : 45
# Proof object formula steps           : 19
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 22
# Proof object simplifying inferences  : 16
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 36
# Processed clauses                    : 365
# ...of these trivial                  : 6
# ...subsumed                          : 204
# ...remaining for further processing  : 154
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 9
# Backward-rewritten                   : 15
# Generated clauses                    : 606
# ...of the previous two non-trivial   : 519
# Contextual simplify-reflections      : 1
# Paramodulations                      : 583
# Factorizations                       : 2
# Equation resolutions                 : 21
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 91
#    Positive orientable unit clauses  : 16
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 12
#    Non-unit-clauses                  : 61
# Current number of unprocessed clauses: 200
# ...number of literals in the above   : 613
# Current number of archived formulas  : 0
# Current number of archived clauses   : 62
# Clause-clause subsumption calls (NU) : 907
# Rec. Clause-clause subsumption calls : 620
# Non-unit clause-clause subsumptions  : 113
# Unit Clause-clause subsumption calls : 495
# Rewrite failures with RHS unbound    : 32
# BW rewrite match attempts            : 45
# BW rewrite match successes           : 10
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6532

# -------------------------------------------------
# User time                : 0.014 s
# System time              : 0.000 s
# Total time               : 0.014 s
# Maximum resident set size: 28592 pages
