# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t125_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(k2_relat_1(X2),X1)=>k8_relat_1(X1,X2)=X2)), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t125_relat_1)).
fof(t71_relat_1, axiom, ![X1]:(k1_relat_1(k6_relat_1(X1))=X1&k2_relat_1(k6_relat_1(X1))=X1), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t71_relat_1)).
fof(fc24_relat_1, axiom, ![X1]:(v1_relat_1(k6_relat_1(X1))&(v4_relat_1(k6_relat_1(X1),X1)&v5_relat_1(k6_relat_1(X1),X1))), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', fc24_relat_1)).
fof(t54_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>k4_relat_1(k5_relat_1(X1,X2))=k5_relat_1(k4_relat_1(X2),k4_relat_1(X1)))), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t54_relat_1)).
fof(t72_relat_1, axiom, ![X1]:k4_relat_1(k6_relat_1(X1))=k6_relat_1(X1), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t72_relat_1)).
fof(t123_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k8_relat_1(X1,X2)=k5_relat_1(X2,k6_relat_1(X1))), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t123_relat_1)).
fof(t76_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(k5_relat_1(X2,k6_relat_1(X1)),X2)&r1_tarski(k5_relat_1(k6_relat_1(X1),X2),X2))), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t76_relat_1)).
fof(t94_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k7_relat_1(X2,X1)=k5_relat_1(k6_relat_1(X1),X2)), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t94_relat_1)).
fof(dt_k5_relat_1, axiom, ![X1, X2]:((v1_relat_1(X1)&v1_relat_1(X2))=>v1_relat_1(k5_relat_1(X1,X2))), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', dt_k5_relat_1)).
fof(t103_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>k7_relat_1(k7_relat_1(X3,X2),X1)=k7_relat_1(X3,X1))), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t103_relat_1)).
fof(t98_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k7_relat_1(X1,k1_relat_1(X1))=X1), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t98_relat_1)).
fof(t25_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>(r1_tarski(k1_relat_1(X1),k1_relat_1(X2))&r1_tarski(k2_relat_1(X1),k2_relat_1(X2)))))), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t25_relat_1)).
fof(t209_relat_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>X2=k7_relat_1(X2,X1)), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t209_relat_1)).
fof(t140_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>k7_relat_1(k8_relat_1(X1,X3),X2)=k8_relat_1(X1,k7_relat_1(X3,X2))), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t140_relat_1)).
fof(t87_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k1_relat_1(k7_relat_1(X2,X1)),X1)), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t87_relat_1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t12_setfam_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t70_enumset1)).
fof(t100_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>k7_relat_1(k7_relat_1(X3,X1),X2)=k7_relat_1(X3,k3_xboole_0(X1,X2))), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t100_relat_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t73_enumset1)).
fof(t43_funct_1, conjecture, ![X1, X2]:k5_relat_1(k6_relat_1(X2),k6_relat_1(X1))=k6_relat_1(k3_xboole_0(X1,X2)), file('data/MPTP/proofs/15f1135b95e747648d38fa58e101f38d.E_input', t43_funct_1)).
fof(c_0_22, plain, ![X6, X7]:(~v1_relat_1(X7)|(~r1_tarski(k2_relat_1(X7),X6)|k8_relat_1(X6,X7)=X7)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t125_relat_1])])).
fof(c_0_23, plain, ![X8]:(k1_relat_1(k6_relat_1(X8))=X8&k2_relat_1(k6_relat_1(X8))=X8), inference(variable_rename,[status(thm)],[t71_relat_1])).
fof(c_0_24, plain, ![X9]:(v1_relat_1(k6_relat_1(X9))&(v4_relat_1(k6_relat_1(X9),X9)&v5_relat_1(k6_relat_1(X9),X9))), inference(variable_rename,[status(thm)],[fc24_relat_1])).
fof(c_0_25, plain, ![X10, X11]:(~v1_relat_1(X10)|(~v1_relat_1(X11)|k4_relat_1(k5_relat_1(X10,X11))=k5_relat_1(k4_relat_1(X11),k4_relat_1(X10)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t54_relat_1])])])).
fof(c_0_26, plain, ![X12]:k4_relat_1(k6_relat_1(X12))=k6_relat_1(X12), inference(variable_rename,[status(thm)],[t72_relat_1])).
fof(c_0_27, plain, ![X13, X14]:(~v1_relat_1(X14)|k8_relat_1(X13,X14)=k5_relat_1(X14,k6_relat_1(X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t123_relat_1])])).
cnf(c_0_28, plain, (k8_relat_1(X2,X1)=X1|~v1_relat_1(X1)|~r1_tarski(k2_relat_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, plain, (k2_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30, plain, (v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_31, plain, ![X15, X16]:((r1_tarski(k5_relat_1(X16,k6_relat_1(X15)),X16)|~v1_relat_1(X16))&(r1_tarski(k5_relat_1(k6_relat_1(X15),X16),X16)|~v1_relat_1(X16))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t76_relat_1])])])).
fof(c_0_32, plain, ![X17, X18]:(~v1_relat_1(X18)|k7_relat_1(X18,X17)=k5_relat_1(k6_relat_1(X17),X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t94_relat_1])])).
fof(c_0_33, plain, ![X19, X20]:(~v1_relat_1(X19)|~v1_relat_1(X20)|v1_relat_1(k5_relat_1(X19,X20))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_relat_1])])).
cnf(c_0_34, plain, (k4_relat_1(k5_relat_1(X1,X2))=k5_relat_1(k4_relat_1(X2),k4_relat_1(X1))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, plain, (k4_relat_1(k6_relat_1(X1))=k6_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (k8_relat_1(X2,X1)=k5_relat_1(X1,k6_relat_1(X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (k8_relat_1(X1,k6_relat_1(X2))=k6_relat_1(X2)|~r1_tarski(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])])).
cnf(c_0_38, plain, (r1_tarski(k5_relat_1(k6_relat_1(X1),X2),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_39, plain, (k7_relat_1(X1,X2)=k5_relat_1(k6_relat_1(X2),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_40, plain, ![X25, X26, X27]:(~v1_relat_1(X27)|(~r1_tarski(X25,X26)|k7_relat_1(k7_relat_1(X27,X26),X25)=k7_relat_1(X27,X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t103_relat_1])])).
cnf(c_0_41, plain, (v1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, plain, (k4_relat_1(k5_relat_1(k6_relat_1(X1),X2))=k5_relat_1(k4_relat_1(X2),k6_relat_1(X1))|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_30])])).
cnf(c_0_43, plain, (k5_relat_1(k6_relat_1(X1),k6_relat_1(X2))=k6_relat_1(X1)|~r1_tarski(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_30])])).
cnf(c_0_44, plain, (r1_tarski(k7_relat_1(X1,X2),X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_45, plain, (k7_relat_1(k7_relat_1(X1,X3),X2)=k7_relat_1(X1,X2)|~v1_relat_1(X1)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_46, plain, (v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_39]), c_0_30])])).
fof(c_0_47, plain, ![X43]:(~v1_relat_1(X43)|k7_relat_1(X43,k1_relat_1(X43))=X43), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t98_relat_1])])).
fof(c_0_48, plain, ![X44, X45]:((r1_tarski(k1_relat_1(X44),k1_relat_1(X45))|~r1_tarski(X44,X45)|~v1_relat_1(X45)|~v1_relat_1(X44))&(r1_tarski(k2_relat_1(X44),k2_relat_1(X45))|~r1_tarski(X44,X45)|~v1_relat_1(X45)|~v1_relat_1(X44))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t25_relat_1])])])])).
fof(c_0_49, plain, ![X46, X47]:(~v1_relat_1(X47)|~v4_relat_1(X47,X46)|X47=k7_relat_1(X47,X46)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t209_relat_1])])).
cnf(c_0_50, plain, (k5_relat_1(k6_relat_1(X1),k6_relat_1(X2))=k6_relat_1(X2)|~r1_tarski(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_35]), c_0_35]), c_0_30])])).
cnf(c_0_51, plain, (r1_tarski(k7_relat_1(X1,X2),k7_relat_1(X1,X3))|~r1_tarski(X2,X3)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])).
cnf(c_0_52, plain, (k7_relat_1(X1,k1_relat_1(X1))=X1|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_53, plain, (r1_tarski(k1_relat_1(X1),k1_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_54, plain, (k1_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_55, plain, ![X48, X49, X50]:(~v1_relat_1(X50)|k7_relat_1(k8_relat_1(X48,X50),X49)=k8_relat_1(X48,k7_relat_1(X50,X49))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t140_relat_1])])).
cnf(c_0_56, plain, (X1=k7_relat_1(X1,X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_57, plain, (v4_relat_1(k6_relat_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_58, plain, ![X51, X52]:(~v1_relat_1(X52)|r1_tarski(k1_relat_1(k7_relat_1(X52,X51)),X51)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t87_relat_1])])).
cnf(c_0_59, plain, (k6_relat_1(X1)=k6_relat_1(X2)|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_43, c_0_50])).
cnf(c_0_60, plain, (r1_tarski(X1,k7_relat_1(X1,X2))|~r1_tarski(k1_relat_1(X1),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_61, plain, (r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(X1,k6_relat_1(X2))|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_30])])).
cnf(c_0_62, plain, (k5_relat_1(k4_relat_1(X1),k6_relat_1(X2))=k4_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_42, c_0_39])).
cnf(c_0_63, plain, (k4_relat_1(k5_relat_1(X1,k6_relat_1(X2)))=k5_relat_1(k6_relat_1(X2),k4_relat_1(X1))|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_30])])).
cnf(c_0_64, plain, (k7_relat_1(k8_relat_1(X2,X1),X3)=k8_relat_1(X2,k7_relat_1(X1,X3))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_65, plain, (k7_relat_1(k6_relat_1(X1),X1)=k6_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_30])])).
fof(c_0_66, plain, ![X21, X22]:k1_setfam_1(k2_tarski(X21,X22))=k3_xboole_0(X21,X22), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_67, plain, ![X23, X24]:k1_enumset1(X23,X23,X24)=k2_tarski(X23,X24), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_68, plain, (r1_tarski(k1_relat_1(k7_relat_1(X1,X2)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_69, plain, (k6_relat_1(k7_relat_1(X1,X2))=k6_relat_1(X1)|~r1_tarski(X1,k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_59, c_0_44])).
cnf(c_0_70, plain, (r1_tarski(X1,k7_relat_1(X1,X2))|~r1_tarski(X1,k6_relat_1(X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_71, plain, (v1_relat_1(k4_relat_1(k7_relat_1(X1,X2)))|~v1_relat_1(k4_relat_1(X1))|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_62]), c_0_30])])).
cnf(c_0_72, plain, (k4_relat_1(k7_relat_1(k6_relat_1(X1),X2))=k5_relat_1(k6_relat_1(X1),k6_relat_1(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_39]), c_0_35]), c_0_30]), c_0_30])])).
cnf(c_0_73, plain, (k7_relat_1(k8_relat_1(X1,k6_relat_1(X2)),X2)=k8_relat_1(X1,k6_relat_1(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_30])])).
cnf(c_0_74, plain, (k7_relat_1(k8_relat_1(X1,X2),X3)=k5_relat_1(k7_relat_1(X2,X3),k6_relat_1(X1))|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_64]), c_0_46])).
fof(c_0_75, plain, ![X28, X29, X30]:(~v1_relat_1(X30)|k7_relat_1(k7_relat_1(X30,X28),X29)=k7_relat_1(X30,k3_xboole_0(X28,X29))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t100_relat_1])])).
cnf(c_0_76, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_77, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
fof(c_0_78, plain, ![X31, X32, X33]:k2_enumset1(X31,X31,X32,X33)=k1_enumset1(X31,X32,X33), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_79, plain, ![X34, X35, X36, X37]:k3_enumset1(X34,X34,X35,X36,X37)=k2_enumset1(X34,X35,X36,X37), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_80, plain, ![X38, X39, X40, X41, X42]:k4_enumset1(X38,X38,X39,X40,X41,X42)=k3_enumset1(X38,X39,X40,X41,X42), inference(variable_rename,[status(thm)],[t73_enumset1])).
cnf(c_0_81, plain, (r1_tarski(X1,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_65]), c_0_54]), c_0_30])])).
cnf(c_0_82, plain, (k6_relat_1(k7_relat_1(X1,X2))=k6_relat_1(X1)|~r1_tarski(X1,k6_relat_1(X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_83, plain, (v1_relat_1(k5_relat_1(k6_relat_1(X1),k6_relat_1(X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_35]), c_0_30]), c_0_30])])).
cnf(c_0_84, plain, (k8_relat_1(X1,k6_relat_1(X2))=k5_relat_1(k6_relat_1(X2),k6_relat_1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_65]), c_0_30])])).
cnf(c_0_85, plain, (k7_relat_1(k7_relat_1(X1,X2),X3)=k7_relat_1(X1,k3_xboole_0(X2,X3))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_86, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_87, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_88, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_89, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_90, plain, (k8_relat_1(k2_relat_1(X1),X1)=X1|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_28, c_0_81])).
cnf(c_0_91, plain, (k6_relat_1(k7_relat_1(k5_relat_1(k6_relat_1(X1),k6_relat_1(X2)),X2))=k6_relat_1(k5_relat_1(k6_relat_1(X1),k6_relat_1(X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_38]), c_0_83]), c_0_30])])).
cnf(c_0_92, plain, (k5_relat_1(k7_relat_1(k6_relat_1(X1),X2),k6_relat_1(X3))=k7_relat_1(k5_relat_1(k6_relat_1(X1),k6_relat_1(X3)),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_84]), c_0_30])])).
cnf(c_0_93, plain, (k7_relat_1(k7_relat_1(X1,X2),X3)=k7_relat_1(X1,k1_setfam_1(k4_enumset1(X2,X2,X2,X2,X2,X3)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_86]), c_0_87]), c_0_88]), c_0_89])).
cnf(c_0_94, plain, (k7_relat_1(k6_relat_1(X1),X2)=k6_relat_1(X2)|~r1_tarski(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_43]), c_0_30])])).
cnf(c_0_95, plain, (k8_relat_1(X1,k6_relat_1(X1))=k6_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_29]), c_0_30])])).
cnf(c_0_96, plain, (k7_relat_1(k5_relat_1(k6_relat_1(X1),k6_relat_1(X2)),X2)=k5_relat_1(k6_relat_1(X1),k6_relat_1(X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_91]), c_0_29])).
cnf(c_0_97, plain, (k7_relat_1(k5_relat_1(k6_relat_1(X1),k6_relat_1(X2)),X3)=k7_relat_1(k6_relat_1(X1),X3)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_37]), c_0_30])]), c_0_92])).
cnf(c_0_98, plain, (v1_relat_1(k7_relat_1(k7_relat_1(X1,X2),X3))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_46, c_0_93])).
cnf(c_0_99, plain, (k7_relat_1(k5_relat_1(k6_relat_1(X1),k6_relat_1(X2)),X3)=k5_relat_1(k6_relat_1(X3),k6_relat_1(X2))|~r1_tarski(X3,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_94]), c_0_30])]), c_0_84]), c_0_84])).
cnf(c_0_100, plain, (k5_relat_1(k6_relat_1(X1),k6_relat_1(X1))=k6_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_95]), c_0_30])])).
cnf(c_0_101, plain, (k5_relat_1(k6_relat_1(X1),k6_relat_1(X2))=k7_relat_1(k6_relat_1(X1),X2)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_102, plain, (v1_relat_1(k7_relat_1(k6_relat_1(X1),X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_65]), c_0_30])])).
cnf(c_0_103, plain, (k6_relat_1(k5_relat_1(k6_relat_1(X1),k6_relat_1(X2)))=k6_relat_1(k6_relat_1(X2))|~r1_tarski(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_99]), c_0_100])).
cnf(c_0_104, plain, (k5_relat_1(k6_relat_1(X1),k6_relat_1(X2))=k5_relat_1(k6_relat_1(X2),k6_relat_1(X1))|~r1_tarski(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_101]), c_0_72]), c_0_35]), c_0_30])])).
cnf(c_0_105, plain, (k6_relat_1(k7_relat_1(k7_relat_1(k6_relat_1(X1),X2),X1))=k6_relat_1(k7_relat_1(k6_relat_1(X1),X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_44]), c_0_102]), c_0_30])])).
cnf(c_0_106, plain, (k6_relat_1(k5_relat_1(k6_relat_1(X1),k6_relat_1(X2)))=k6_relat_1(k6_relat_1(X1))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_107, plain, (k7_relat_1(k7_relat_1(k6_relat_1(X1),X2),X1)=k7_relat_1(k6_relat_1(X1),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_105]), c_0_29])).
cnf(c_0_108, plain, (k6_relat_1(k7_relat_1(k6_relat_1(X1),X2))=k6_relat_1(k6_relat_1(X2))|~r1_tarski(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_39]), c_0_30])])).
cnf(c_0_109, plain, (k7_relat_1(X1,k1_relat_1(k7_relat_1(X1,X2)))=k7_relat_1(X1,X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_45]), c_0_46]), c_0_68])).
cnf(c_0_110, plain, (r1_tarski(k1_relat_1(k7_relat_1(k6_relat_1(X1),X2)),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_107]), c_0_102])])).
fof(c_0_111, negated_conjecture, ~(![X1, X2]:k5_relat_1(k6_relat_1(X2),k6_relat_1(X1))=k6_relat_1(k3_xboole_0(X1,X2))), inference(assume_negation,[status(cth)],[t43_funct_1])).
cnf(c_0_112, plain, (k6_relat_1(k6_relat_1(k1_relat_1(k7_relat_1(k6_relat_1(X1),X2))))=k6_relat_1(k7_relat_1(k6_relat_1(X1),X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_110]), c_0_30])])).
fof(c_0_113, negated_conjecture, k5_relat_1(k6_relat_1(esk2_0),k6_relat_1(esk1_0))!=k6_relat_1(k3_xboole_0(esk1_0,esk2_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_111])])])).
cnf(c_0_114, plain, (k6_relat_1(k1_relat_1(k7_relat_1(k6_relat_1(X1),X2)))=k7_relat_1(k6_relat_1(X1),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_112]), c_0_29])).
cnf(c_0_115, negated_conjecture, (k5_relat_1(k6_relat_1(esk2_0),k6_relat_1(esk1_0))!=k6_relat_1(k3_xboole_0(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_116, plain, (k5_relat_1(k6_relat_1(X1),k6_relat_1(X2))=k7_relat_1(k6_relat_1(X1),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_114]), c_0_72])).
cnf(c_0_117, negated_conjecture, (k5_relat_1(k6_relat_1(esk2_0),k6_relat_1(esk1_0))!=k6_relat_1(k1_setfam_1(k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_115, c_0_86]), c_0_87]), c_0_88]), c_0_89])).
cnf(c_0_118, plain, (k7_relat_1(k7_relat_1(k6_relat_1(k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2))),X1),X2)=k6_relat_1(k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_93]), c_0_30])])).
cnf(c_0_119, plain, (k7_relat_1(k6_relat_1(X1),X2)=k7_relat_1(k6_relat_1(X2),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_116]), c_0_72]), c_0_116]), c_0_35]), c_0_116]), c_0_30])])).
cnf(c_0_120, plain, (k7_relat_1(k5_relat_1(k6_relat_1(X1),k6_relat_1(X2)),X1)=k5_relat_1(k6_relat_1(X1),k6_relat_1(X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_84]), c_0_84])).
cnf(c_0_121, negated_conjecture, (k6_relat_1(k1_setfam_1(k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)))!=k7_relat_1(k6_relat_1(esk2_0),esk1_0)), inference(rw,[status(thm)],[c_0_117, c_0_116])).
cnf(c_0_122, plain, (k7_relat_1(k7_relat_1(k6_relat_1(X1),k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2))),X2)=k6_relat_1(k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2)))), inference(rw,[status(thm)],[c_0_118, c_0_119])).
cnf(c_0_123, plain, (k7_relat_1(k7_relat_1(k6_relat_1(X1),X2),X2)=k7_relat_1(k6_relat_1(X1),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_39]), c_0_30])])).
cnf(c_0_124, negated_conjecture, (k6_relat_1(k1_setfam_1(k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)))!=k7_relat_1(k6_relat_1(esk1_0),esk2_0)), inference(rw,[status(thm)],[c_0_121, c_0_119])).
cnf(c_0_125, plain, (k6_relat_1(k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2)))=k7_relat_1(k6_relat_1(X1),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_93]), c_0_65]), c_0_123]), c_0_30])])).
cnf(c_0_126, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_124, c_0_125])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 127
# Proof object clause steps            : 82
# Proof object formula steps           : 45
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 22
# Proof object generating inferences   : 50
# Proof object simplifying inferences  : 109
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 22
# Processed clauses                    : 9984
# ...of these trivial                  : 196
# ...subsumed                          : 8311
# ...remaining for further processing  : 1477
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 22
# Backward-rewritten                   : 120
# Generated clauses                    : 279424
# ...of the previous two non-trivial   : 254414
# Contextual simplify-reflections      : 130
# Paramodulations                      : 279424
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1313
#    Positive orientable unit clauses  : 89
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 1223
# Current number of unprocessed clauses: 243899
# ...number of literals in the above   : 902509
# Current number of archived formulas  : 0
# Current number of archived clauses   : 169
# Clause-clause subsumption calls (NU) : 197645
# Rec. Clause-clause subsumption calls : 153762
# Non-unit clause-clause subsumptions  : 8436
# Unit Clause-clause subsumption calls : 257
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 386
# BW rewrite match successes           : 76
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5960428

# -------------------------------------------------
# User time                : 2.857 s
# System time              : 0.066 s
# Total time               : 2.923 s
# Maximum resident set size: 180112 pages
