# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t122_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k3_xboole_0(X1,X2),X3)=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&k2_zfmisc_1(X3,k3_xboole_0(X1,X2))=k3_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))), file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', t122_zfmisc_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', t7_xboole_1)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', t1_boole)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', t28_xboole_1)).
fof(t118_zfmisc_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>(r1_tarski(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&r1_tarski(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2)))), file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', t118_zfmisc_1)).
fof(t106_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k4_xboole_0(X2,X3))=>(r1_tarski(X1,X2)&r1_xboole_0(X1,X3))), file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', t106_xboole_1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', t4_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', commutativity_k3_xboole_0)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', d7_xboole_0)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', t113_zfmisc_1)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', t2_boole)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', t65_xboole_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', t3_xboole_0)).
fof(t69_xboole_1, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>~((r1_tarski(X2,X1)&r1_xboole_0(X2,X1)))), file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', t69_xboole_1)).
fof(t117_zfmisc_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&((r1_tarski(k2_zfmisc_1(X2,X1),k2_zfmisc_1(X3,X1))|r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))&~(r1_tarski(X2,X3))))), file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', t117_zfmisc_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', t2_xboole_1)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', l13_xboole_0)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', t3_boole)).
fof(t138_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=>(k2_zfmisc_1(X1,X2)=k1_xboole_0|(r1_tarski(X1,X3)&r1_tarski(X2,X4)))), file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', t138_zfmisc_1)).
fof(t123_zfmisc_1, axiom, ![X1, X2, X3, X4]:k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)), file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', t123_zfmisc_1)).
fof(t63_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_xboole_0(X2,X3))=>r1_xboole_0(X1,X3)), file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', t63_xboole_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/1650948e4e414a918f0961317d8a9c76.E_input', l32_xboole_1)).
fof(c_0_22, plain, ![X9, X10, X11]:(k2_zfmisc_1(k3_xboole_0(X9,X10),X11)=k3_xboole_0(k2_zfmisc_1(X9,X11),k2_zfmisc_1(X10,X11))&k2_zfmisc_1(X11,k3_xboole_0(X9,X10))=k3_xboole_0(k2_zfmisc_1(X11,X9),k2_zfmisc_1(X11,X10))), inference(variable_rename,[status(thm)],[t122_zfmisc_1])).
fof(c_0_23, plain, ![X40, X41]:r1_tarski(X40,k2_xboole_0(X40,X41)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_24, plain, ![X19]:k2_xboole_0(X19,k1_xboole_0)=X19, inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_25, plain, (k2_zfmisc_1(X1,k3_xboole_0(X2,X3))=k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),X3)=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_27, plain, ![X51, X52]:(~r1_tarski(X51,X52)|k3_xboole_0(X51,X52)=X51), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_28, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_30, plain, ![X14, X15, X16]:((r1_tarski(k2_zfmisc_1(X14,X16),k2_zfmisc_1(X15,X16))|~r1_tarski(X14,X15))&(r1_tarski(k2_zfmisc_1(X16,X14),k2_zfmisc_1(X16,X15))|~r1_tarski(X14,X15))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t118_zfmisc_1])])])).
cnf(c_0_31, plain, (k2_zfmisc_1(k3_xboole_0(X1,X1),X2)=k2_zfmisc_1(X1,k3_xboole_0(X2,X2))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_32, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
fof(c_0_34, plain, ![X26, X27, X28]:((r1_tarski(X26,X27)|~r1_tarski(X26,k4_xboole_0(X27,X28)))&(r1_xboole_0(X26,X28)|~r1_tarski(X26,k4_xboole_0(X27,X28)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t106_xboole_1])])])).
cnf(c_0_35, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_36, plain, (k2_zfmisc_1(X1,k3_xboole_0(X2,X2))=k2_zfmisc_1(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])])).
cnf(c_0_37, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k4_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_38, plain, ![X45, X46, X48, X49, X50]:((r1_xboole_0(X45,X46)|r2_hidden(esk6_2(X45,X46),k3_xboole_0(X45,X46)))&(~r2_hidden(X50,k3_xboole_0(X48,X49))|~r1_xboole_0(X48,X49))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_39, plain, ![X17, X18]:k3_xboole_0(X17,X18)=k3_xboole_0(X18,X17), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_40, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))|~r1_tarski(X2,k3_xboole_0(X3,X3))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_41, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_37, c_0_33])).
cnf(c_0_42, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_43, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_44, plain, (r1_tarski(k2_zfmisc_1(X1,k4_xboole_0(k3_xboole_0(X2,X2),X3)),k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
fof(c_0_45, plain, ![X38, X39]:((~r1_xboole_0(X38,X39)|k3_xboole_0(X38,X39)=k1_xboole_0)&(k3_xboole_0(X38,X39)!=k1_xboole_0|r1_xboole_0(X38,X39))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_46, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k4_xboole_0(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_47, plain, ![X24, X25]:((k2_zfmisc_1(X24,X25)!=k1_xboole_0|(X24=k1_xboole_0|X25=k1_xboole_0))&((X24!=k1_xboole_0|k2_zfmisc_1(X24,X25)=k1_xboole_0)&(X25!=k1_xboole_0|k2_zfmisc_1(X24,X25)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_48, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))|~r1_tarski(k3_xboole_0(X2,X2),X3)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
fof(c_0_49, plain, ![X12]:k3_xboole_0(X12,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_50, plain, ![X54]:r1_xboole_0(X54,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_51, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_52, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk6_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_53, plain, (k2_zfmisc_1(X1,k3_xboole_0(X2,X3))=k2_zfmisc_1(X1,X2)|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(spm,[status(thm)],[c_0_25, c_0_32])).
cnf(c_0_54, plain, (r1_tarski(k2_zfmisc_1(X1,k4_xboole_0(X2,X3)),k2_zfmisc_1(X1,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_32]), c_0_33])])).
cnf(c_0_55, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_56, plain, (r1_xboole_0(k4_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_46, c_0_33])).
cnf(c_0_57, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_58, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,k2_xboole_0(k3_xboole_0(X2,X2),X3)))), inference(spm,[status(thm)],[c_0_48, c_0_28])).
cnf(c_0_59, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_60, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_50])).
fof(c_0_61, plain, ![X32, X33, X35, X36, X37]:(((r2_hidden(esk5_2(X32,X33),X32)|r1_xboole_0(X32,X33))&(r2_hidden(esk5_2(X32,X33),X33)|r1_xboole_0(X32,X33)))&(~r2_hidden(X37,X35)|~r2_hidden(X37,X36)|~r1_xboole_0(X35,X36))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_62, plain, ![X58, X59]:(v1_xboole_0(X59)|(~r1_tarski(X59,X58)|~r1_xboole_0(X59,X58))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t69_xboole_1])])])).
cnf(c_0_63, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
fof(c_0_64, plain, ![X29, X30, X31]:((~r1_tarski(k2_zfmisc_1(X30,X29),k2_zfmisc_1(X31,X29))|r1_tarski(X30,X31)|X29=k1_xboole_0)&(~r1_tarski(k2_zfmisc_1(X29,X30),k2_zfmisc_1(X29,X31))|r1_tarski(X30,X31)|X29=k1_xboole_0)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t117_zfmisc_1])])])])).
cnf(c_0_65, plain, (k2_zfmisc_1(X1,k3_xboole_0(X2,k4_xboole_0(X2,X3)))=k2_zfmisc_1(X1,k4_xboole_0(X2,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_43])).
cnf(c_0_66, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X1))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_43])).
cnf(c_0_67, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_57])).
fof(c_0_68, plain, ![X13]:r1_tarski(k1_xboole_0,X13), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_69, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,k2_xboole_0(X2,X3)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_32]), c_0_33])])).
cnf(c_0_70, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_59]), c_0_60])])).
cnf(c_0_71, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_72, plain, (v1_xboole_0(X1)|~r1_tarski(X1,X2)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_73, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_63, c_0_56])).
cnf(c_0_74, plain, (r1_tarski(X2,X3)|X1=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_75, plain, (k2_zfmisc_1(X1,k4_xboole_0(X2,X2))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67])).
cnf(c_0_76, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_77, plain, (k2_zfmisc_1(X1,k3_xboole_0(X2,k2_xboole_0(X2,X3)))=k2_zfmisc_1(X1,X2)), inference(spm,[status(thm)],[c_0_53, c_0_69])).
cnf(c_0_78, plain, (r1_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_79, plain, (v1_xboole_0(X1)|~r1_tarski(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_80, plain, (X1=k1_xboole_0|r1_tarski(k4_xboole_0(X2,X2),X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76])])).
cnf(c_0_81, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,k2_xboole_0(k3_xboole_0(X2,k2_xboole_0(X2,X3)),X4)))), inference(spm,[status(thm)],[c_0_69, c_0_77])).
cnf(c_0_82, plain, (v1_xboole_0(k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_78]), c_0_76])])).
cnf(c_0_83, plain, (X1=k1_xboole_0|v1_xboole_0(k4_xboole_0(X2,X2))), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_84, plain, (X1=k1_xboole_0|r1_tarski(X2,k2_xboole_0(k3_xboole_0(X2,k2_xboole_0(X2,X3)),X4))), inference(spm,[status(thm)],[c_0_74, c_0_81])).
fof(c_0_85, plain, ![X53]:(~v1_xboole_0(X53)|X53=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_86, plain, (v1_xboole_0(k4_xboole_0(X1,X1))|v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_87, plain, (X1=k1_xboole_0|r1_tarski(X2,k3_xboole_0(X2,k2_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_84, c_0_29])).
cnf(c_0_88, plain, (r1_tarski(k2_zfmisc_1(X1,k4_xboole_0(k3_xboole_0(X2,k2_xboole_0(X2,X3)),X4)),k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_54, c_0_77])).
fof(c_0_89, plain, ![X44]:k4_xboole_0(X44,k1_xboole_0)=X44, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_90, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_91, plain, (v1_xboole_0(k4_xboole_0(X1,X1))), inference(ef,[status(thm)],[c_0_86])).
cnf(c_0_92, plain, (r1_tarski(X1,k3_xboole_0(X1,k2_xboole_0(X1,X2)))|r1_tarski(X3,X4)), inference(spm,[status(thm)],[c_0_76, c_0_87])).
cnf(c_0_93, plain, (r1_xboole_0(k3_xboole_0(X1,X2),X3)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_51, c_0_71])).
cnf(c_0_94, plain, (X1=k1_xboole_0|r1_tarski(k4_xboole_0(k3_xboole_0(X2,k2_xboole_0(X2,X3)),X4),X2)), inference(spm,[status(thm)],[c_0_74, c_0_88])).
cnf(c_0_95, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_89])).
fof(c_0_96, negated_conjecture, ~(![X1, X2, X3, X4]:(r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=>(k2_zfmisc_1(X1,X2)=k1_xboole_0|(r1_tarski(X1,X3)&r1_tarski(X2,X4))))), inference(assume_negation,[status(cth)],[t138_zfmisc_1])).
fof(c_0_97, plain, ![X20, X21, X22, X23]:k2_zfmisc_1(k3_xboole_0(X20,X21),k3_xboole_0(X22,X23))=k3_xboole_0(k2_zfmisc_1(X20,X22),k2_zfmisc_1(X21,X23)), inference(variable_rename,[status(thm)],[t123_zfmisc_1])).
cnf(c_0_98, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_99, plain, (r1_tarski(X1,k3_xboole_0(X1,k2_xboole_0(X1,X2)))), inference(ef,[status(thm)],[c_0_92])).
cnf(c_0_100, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k1_xboole_0|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_55, c_0_93])).
cnf(c_0_101, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_102, plain, (X1=k1_xboole_0|r1_tarski(k3_xboole_0(X2,k2_xboole_0(X2,X3)),X2)), inference(spm,[status(thm)],[c_0_94, c_0_95])).
fof(c_0_103, plain, ![X55, X56, X57]:(~r1_tarski(X55,X56)|~r1_xboole_0(X56,X57)|r1_xboole_0(X55,X57)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_xboole_1])])).
fof(c_0_104, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))&(k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0&(~r1_tarski(esk1_0,esk3_0)|~r1_tarski(esk2_0,esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_96])])])).
cnf(c_0_105, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_106, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_37, c_0_98])).
cnf(c_0_107, plain, (r1_tarski(k3_xboole_0(X1,X2),k1_xboole_0)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_108, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_101])).
cnf(c_0_109, plain, (k3_xboole_0(X1,X2)=X2|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_43, c_0_32])).
cnf(c_0_110, plain, (r1_tarski(k3_xboole_0(X1,k2_xboole_0(X1,X2)),X1)|r1_tarski(X3,X4)), inference(spm,[status(thm)],[c_0_76, c_0_102])).
cnf(c_0_111, plain, (r1_xboole_0(X1,X3)|~r1_tarski(X1,X2)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_112, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_113, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,k4_xboole_0(X4,X2)))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_66]), c_0_67])).
cnf(c_0_114, plain, (r1_tarski(k3_xboole_0(X1,X2),X3)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_106, c_0_107])).
cnf(c_0_115, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(k4_xboole_0(X3,X1),X4))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_66]), c_0_108])).
cnf(c_0_116, plain, (X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_32, c_0_109])).
cnf(c_0_117, plain, (r1_tarski(k3_xboole_0(X1,k2_xboole_0(X1,X2)),X1)), inference(ef,[status(thm)],[c_0_110])).
cnf(c_0_118, plain, (~r2_hidden(X1,k2_zfmisc_1(X2,k3_xboole_0(X3,X4)))|~r1_xboole_0(k2_zfmisc_1(X2,X3),k2_zfmisc_1(X2,X4))), inference(spm,[status(thm)],[c_0_42, c_0_25])).
cnf(c_0_119, plain, (r2_hidden(esk5_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_120, negated_conjecture, (r1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0),X1)|~r1_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),X1)), inference(spm,[status(thm)],[c_0_111, c_0_112])).
cnf(c_0_121, plain, (r1_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,k4_xboole_0(X4,X2)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_113]), c_0_70])).
cnf(c_0_122, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_79, c_0_114])).
cnf(c_0_123, plain, (r1_xboole_0(k3_xboole_0(X1,X2),X3)|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_42, c_0_71])).
cnf(c_0_124, plain, (r1_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(k4_xboole_0(X3,X1),X4))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_115]), c_0_70])).
cnf(c_0_125, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_99])])).
cnf(c_0_126, plain, (r1_xboole_0(X1,k2_zfmisc_1(X2,k3_xboole_0(X3,X4)))|~r1_xboole_0(k2_zfmisc_1(X2,X3),k2_zfmisc_1(X2,X4))), inference(spm,[status(thm)],[c_0_118, c_0_119])).
cnf(c_0_127, negated_conjecture, (r1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(X1,k4_xboole_0(X2,esk4_0)))), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_128, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_122, c_0_43])).
cnf(c_0_129, plain, (r1_xboole_0(k2_zfmisc_1(k3_xboole_0(X1,X2),X3),X4)|~r1_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_123, c_0_26])).
cnf(c_0_130, negated_conjecture, (r1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(k4_xboole_0(X1,esk3_0),X2))), inference(spm,[status(thm)],[c_0_120, c_0_124])).
cnf(c_0_131, plain, (v1_xboole_0(X1)|~r1_xboole_0(k2_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_122, c_0_125])).
cnf(c_0_132, negated_conjecture, (r1_xboole_0(X1,k2_zfmisc_1(esk1_0,k3_xboole_0(esk2_0,k4_xboole_0(X2,esk4_0))))), inference(spm,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_133, plain, (v1_xboole_0(X1)|~r1_xboole_0(X1,k2_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_128, c_0_125])).
cnf(c_0_134, negated_conjecture, (r1_xboole_0(k2_zfmisc_1(k3_xboole_0(esk1_0,k4_xboole_0(X1,esk3_0)),esk2_0),X2)), inference(spm,[status(thm)],[c_0_129, c_0_130])).
cnf(c_0_135, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(esk1_0,k3_xboole_0(esk2_0,k4_xboole_0(X1,esk4_0))))), inference(spm,[status(thm)],[c_0_131, c_0_132])).
cnf(c_0_136, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(k3_xboole_0(esk1_0,k4_xboole_0(X1,esk3_0)),esk2_0))), inference(spm,[status(thm)],[c_0_133, c_0_134])).
cnf(c_0_137, negated_conjecture, (k2_zfmisc_1(esk1_0,k3_xboole_0(esk2_0,k4_xboole_0(X1,esk4_0)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_90, c_0_135])).
cnf(c_0_138, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_139, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk1_0,k4_xboole_0(X1,esk3_0)),esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_90, c_0_136])).
cnf(c_0_140, negated_conjecture, (r1_xboole_0(X1,k2_zfmisc_1(esk1_0,k4_xboole_0(esk2_0,esk4_0)))), inference(spm,[status(thm)],[c_0_132, c_0_65])).
cnf(c_0_141, negated_conjecture, (~r1_tarski(esk2_0,k4_xboole_0(X1,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_32]), c_0_138])).
cnf(c_0_142, negated_conjecture, (k3_xboole_0(esk1_0,k4_xboole_0(X1,esk3_0))=k1_xboole_0|r1_tarski(esk2_0,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_139]), c_0_76])])).
cnf(c_0_143, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(esk1_0,k4_xboole_0(esk2_0,esk4_0)))), inference(spm,[status(thm)],[c_0_131, c_0_140])).
cnf(c_0_144, plain, (X1=k1_xboole_0|r1_tarski(k4_xboole_0(k3_xboole_0(X2,X2),X3),X2)), inference(spm,[status(thm)],[c_0_74, c_0_44])).
cnf(c_0_145, plain, (X1=k1_xboole_0|r1_tarski(X2,k2_xboole_0(k3_xboole_0(X2,X2),X3))), inference(spm,[status(thm)],[c_0_74, c_0_58])).
cnf(c_0_146, plain, (r1_xboole_0(k4_xboole_0(k4_xboole_0(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_46, c_0_41])).
cnf(c_0_147, negated_conjecture, (k3_xboole_0(esk1_0,k4_xboole_0(X1,esk3_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_141, c_0_142])).
fof(c_0_148, plain, ![X42, X43]:((k4_xboole_0(X42,X43)!=k1_xboole_0|r1_tarski(X42,X43))&(~r1_tarski(X42,X43)|k4_xboole_0(X42,X43)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_149, plain, (r1_tarski(X1,X3)|X2=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_150, negated_conjecture, (k2_zfmisc_1(esk1_0,k4_xboole_0(esk2_0,esk4_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_90, c_0_143])).
cnf(c_0_151, plain, (X1=k1_xboole_0|r1_tarski(k3_xboole_0(X2,X2),X2)), inference(spm,[status(thm)],[c_0_144, c_0_95])).
cnf(c_0_152, plain, (X1=k1_xboole_0|r1_tarski(X2,k3_xboole_0(X2,X2))), inference(spm,[status(thm)],[c_0_145, c_0_29])).
cnf(c_0_153, plain, (r1_xboole_0(X1,k4_xboole_0(k4_xboole_0(X2,X1),X3))), inference(spm,[status(thm)],[c_0_63, c_0_146])).
cnf(c_0_154, negated_conjecture, (k2_zfmisc_1(X1,k4_xboole_0(esk1_0,esk3_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_147]), c_0_67])).
cnf(c_0_155, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_148])).
cnf(c_0_156, negated_conjecture, (k4_xboole_0(esk2_0,esk4_0)=k1_xboole_0|r1_tarski(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_150]), c_0_76])])).
cnf(c_0_157, plain, (r1_tarski(k3_xboole_0(X1,X1),X1)|r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_76, c_0_151])).
cnf(c_0_158, plain, (r1_tarski(X1,k3_xboole_0(X1,X1))|r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_76, c_0_152])).
cnf(c_0_159, plain, (v1_xboole_0(X1)|~r1_tarski(X1,k4_xboole_0(k4_xboole_0(X2,X1),X3))), inference(spm,[status(thm)],[c_0_72, c_0_153])).
cnf(c_0_160, negated_conjecture, (X1=k1_xboole_0|r1_tarski(k4_xboole_0(esk1_0,esk3_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_154]), c_0_76])])).
cnf(c_0_161, plain, (k4_xboole_0(X1,X2)=X1|~r1_tarski(X1,k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_116, c_0_41])).
cnf(c_0_162, negated_conjecture, (r1_tarski(esk2_0,esk4_0)|r1_tarski(esk1_0,X1)), inference(spm,[status(thm)],[c_0_155, c_0_156])).
cnf(c_0_163, plain, (r1_tarski(k3_xboole_0(X1,X1),X1)), inference(ef,[status(thm)],[c_0_157])).
cnf(c_0_164, plain, (r1_tarski(X1,k3_xboole_0(X1,X1))), inference(ef,[status(thm)],[c_0_158])).
cnf(c_0_165, negated_conjecture, (X1=k1_xboole_0|v1_xboole_0(k4_xboole_0(esk1_0,esk3_0))), inference(spm,[status(thm)],[c_0_159, c_0_160])).
cnf(c_0_166, negated_conjecture, (k4_xboole_0(esk1_0,X1)=esk1_0|r1_tarski(esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_161, c_0_162])).
cnf(c_0_167, plain, (k3_xboole_0(X1,X1)=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_163]), c_0_164])])).
cnf(c_0_168, negated_conjecture, (v1_xboole_0(k4_xboole_0(esk1_0,esk3_0))|v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_82, c_0_165])).
cnf(c_0_169, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0))|r1_tarski(esk2_0,esk4_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_166]), c_0_167])).
cnf(c_0_170, negated_conjecture, (v1_xboole_0(k4_xboole_0(esk1_0,esk3_0))), inference(ef,[status(thm)],[c_0_168])).
cnf(c_0_171, negated_conjecture, (~r1_tarski(esk1_0,esk3_0)|~r1_tarski(esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_172, negated_conjecture, (r1_tarski(esk2_0,esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_169]), c_0_138])).
cnf(c_0_173, negated_conjecture, (k4_xboole_0(esk1_0,esk3_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_90, c_0_170])).
cnf(c_0_174, negated_conjecture, (~r1_tarski(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_171, c_0_172])])).
cnf(c_0_175, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_173]), c_0_174]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 176
# Proof object clause steps            : 131
# Proof object formula steps           : 45
# Proof object conjectures             : 34
# Proof object clause conjectures      : 31
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 30
# Proof object initial formulas used   : 22
# Proof object generating inferences   : 98
# Proof object simplifying inferences  : 38
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 35
# Processed clauses                    : 2602
# ...of these trivial                  : 137
# ...subsumed                          : 1822
# ...remaining for further processing  : 643
# Other redundant clauses eliminated   : 48
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 66
# Backward-rewritten                   : 112
# Generated clauses                    : 33798
# ...of the previous two non-trivial   : 19031
# Contextual simplify-reflections      : 0
# Paramodulations                      : 33726
# Factorizations                       : 24
# Equation resolutions                 : 48
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 426
#    Positive orientable unit clauses  : 125
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 289
# Current number of unprocessed clauses: 14918
# ...number of literals in the above   : 33977
# Current number of archived formulas  : 0
# Current number of archived clauses   : 213
# Clause-clause subsumption calls (NU) : 11161
# Rec. Clause-clause subsumption calls : 10116
# Non-unit clause-clause subsumptions  : 1209
# Unit Clause-clause subsumption calls : 858
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 482
# BW rewrite match successes           : 159
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 339241

# -------------------------------------------------
# User time                : 0.170 s
# System time              : 0.007 s
# Total time               : 0.177 s
# Maximum resident set size: 28752 pages
