# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_mcart_1, axiom, ![X1, X2, X3, X4]:k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4), file('data/MPTP/proofs/1671842233a944e7a19b76ca92158b37.E_input', d4_mcart_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/1671842233a944e7a19b76ca92158b37.E_input', d3_mcart_1)).
fof(t78_mcart_1, axiom, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&?[X6, X7, X8, X9]:(X5=k4_mcart_1(X6,X7,X8,X9)&~((k8_mcart_1(X1,X2,X3,X4,X5)=X6&(k9_mcart_1(X1,X2,X3,X4,X5)=X7&(k10_mcart_1(X1,X2,X3,X4,X5)=X8&k11_mcart_1(X1,X2,X3,X4,X5)=X9))))))))))), file('data/MPTP/proofs/1671842233a944e7a19b76ca92158b37.E_input', t78_mcart_1)).
fof(t61_mcart_1, axiom, ![X1, X2, X3, X4]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&~(![X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>(k8_mcart_1(X1,X2,X3,X4,X5)=k1_mcart_1(k1_mcart_1(k1_mcart_1(X5)))&(k9_mcart_1(X1,X2,X3,X4,X5)=k2_mcart_1(k1_mcart_1(k1_mcart_1(X5)))&(k10_mcart_1(X1,X2,X3,X4,X5)=k2_mcart_1(k1_mcart_1(X5))&k11_mcart_1(X1,X2,X3,X4,X5)=k2_mcart_1(X5))))))))))), file('data/MPTP/proofs/1671842233a944e7a19b76ca92158b37.E_input', t61_mcart_1)).
fof(l68_mcart_1, axiom, ![X1, X2, X3, X4]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&?[X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))&![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>![X9]:(m1_subset_1(X9,X4)=>~(X5=k4_mcart_1(X6,X7,X8,X9)))))))))))), file('data/MPTP/proofs/1671842233a944e7a19b76ca92158b37.E_input', l68_mcart_1)).
fof(t79_mcart_1, conjecture, ![X1, X2, X3, X4, X5, X6]:(m1_subset_1(X6,k4_zfmisc_1(X1,X2,X3,X4))=>(![X7]:(m1_subset_1(X7,X1)=>![X8]:(m1_subset_1(X8,X2)=>![X9]:(m1_subset_1(X9,X3)=>![X10]:(m1_subset_1(X10,X4)=>(X6=k4_mcart_1(X7,X8,X9,X10)=>X5=X7)))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|(X4=k1_xboole_0|X5=k8_mcart_1(X1,X2,X3,X4,X6))))))), file('data/MPTP/proofs/1671842233a944e7a19b76ca92158b37.E_input', t79_mcart_1)).
fof(c_0_6, plain, ![X35, X36, X37, X38]:k4_mcart_1(X35,X36,X37,X38)=k4_tarski(k3_mcart_1(X35,X36,X37),X38), inference(variable_rename,[status(thm)],[d4_mcart_1])).
fof(c_0_7, plain, ![X39, X40, X41]:k3_mcart_1(X39,X40,X41)=k4_tarski(k4_tarski(X39,X40),X41), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_8, plain, ![X42, X43, X44, X45, X46, X47, X48, X49, X50]:((k8_mcart_1(X42,X43,X44,X45,X46)=X47|X46!=k4_mcart_1(X47,X48,X49,X50)|X45=k1_xboole_0|X44=k1_xboole_0|X43=k1_xboole_0|X42=k1_xboole_0|~m1_subset_1(X46,k4_zfmisc_1(X42,X43,X44,X45)))&((k9_mcart_1(X42,X43,X44,X45,X46)=X48|X46!=k4_mcart_1(X47,X48,X49,X50)|X45=k1_xboole_0|X44=k1_xboole_0|X43=k1_xboole_0|X42=k1_xboole_0|~m1_subset_1(X46,k4_zfmisc_1(X42,X43,X44,X45)))&((k10_mcart_1(X42,X43,X44,X45,X46)=X49|X46!=k4_mcart_1(X47,X48,X49,X50)|X45=k1_xboole_0|X44=k1_xboole_0|X43=k1_xboole_0|X42=k1_xboole_0|~m1_subset_1(X46,k4_zfmisc_1(X42,X43,X44,X45)))&(k11_mcart_1(X42,X43,X44,X45,X46)=X50|X46!=k4_mcart_1(X47,X48,X49,X50)|X45=k1_xboole_0|X44=k1_xboole_0|X43=k1_xboole_0|X42=k1_xboole_0|~m1_subset_1(X46,k4_zfmisc_1(X42,X43,X44,X45)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t78_mcart_1])])])])])).
cnf(c_0_9, plain, (k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (k11_mcart_1(X1,X2,X3,X4,X5)=X6|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X5!=k4_mcart_1(X7,X8,X9,X6)|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_9, c_0_10])).
fof(c_0_13, plain, ![X30, X31, X32, X33, X34]:((k8_mcart_1(X30,X31,X32,X33,X34)=k1_mcart_1(k1_mcart_1(k1_mcart_1(X34)))|~m1_subset_1(X34,k4_zfmisc_1(X30,X31,X32,X33))|X33=k1_xboole_0|X32=k1_xboole_0|X31=k1_xboole_0|X30=k1_xboole_0)&((k9_mcart_1(X30,X31,X32,X33,X34)=k2_mcart_1(k1_mcart_1(k1_mcart_1(X34)))|~m1_subset_1(X34,k4_zfmisc_1(X30,X31,X32,X33))|X33=k1_xboole_0|X32=k1_xboole_0|X31=k1_xboole_0|X30=k1_xboole_0)&((k10_mcart_1(X30,X31,X32,X33,X34)=k2_mcart_1(k1_mcart_1(X34))|~m1_subset_1(X34,k4_zfmisc_1(X30,X31,X32,X33))|X33=k1_xboole_0|X32=k1_xboole_0|X31=k1_xboole_0|X30=k1_xboole_0)&(k11_mcart_1(X30,X31,X32,X33,X34)=k2_mcart_1(X34)|~m1_subset_1(X34,k4_zfmisc_1(X30,X31,X32,X33))|X33=k1_xboole_0|X32=k1_xboole_0|X31=k1_xboole_0|X30=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t61_mcart_1])])])])])).
cnf(c_0_14, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k11_mcart_1(X1,X2,X3,X4,X5)=X6|X5!=k4_tarski(k4_tarski(k4_tarski(X7,X8),X9),X6)|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(rw,[status(thm)],[c_0_11, c_0_12])).
fof(c_0_15, plain, ![X21, X22, X23, X24, X25]:((m1_subset_1(esk7_5(X21,X22,X23,X24,X25),X21)|~m1_subset_1(X25,k4_zfmisc_1(X21,X22,X23,X24))|X24=k1_xboole_0|X23=k1_xboole_0|X22=k1_xboole_0|X21=k1_xboole_0)&((m1_subset_1(esk8_5(X21,X22,X23,X24,X25),X22)|~m1_subset_1(X25,k4_zfmisc_1(X21,X22,X23,X24))|X24=k1_xboole_0|X23=k1_xboole_0|X22=k1_xboole_0|X21=k1_xboole_0)&((m1_subset_1(esk9_5(X21,X22,X23,X24,X25),X23)|~m1_subset_1(X25,k4_zfmisc_1(X21,X22,X23,X24))|X24=k1_xboole_0|X23=k1_xboole_0|X22=k1_xboole_0|X21=k1_xboole_0)&((m1_subset_1(esk10_5(X21,X22,X23,X24,X25),X24)|~m1_subset_1(X25,k4_zfmisc_1(X21,X22,X23,X24))|X24=k1_xboole_0|X23=k1_xboole_0|X22=k1_xboole_0|X21=k1_xboole_0)&(X25=k4_mcart_1(esk7_5(X21,X22,X23,X24,X25),esk8_5(X21,X22,X23,X24,X25),esk9_5(X21,X22,X23,X24,X25),esk10_5(X21,X22,X23,X24,X25))|~m1_subset_1(X25,k4_zfmisc_1(X21,X22,X23,X24))|X24=k1_xboole_0|X23=k1_xboole_0|X22=k1_xboole_0|X21=k1_xboole_0))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l68_mcart_1])])])])])])])).
cnf(c_0_16, plain, (k11_mcart_1(X1,X2,X3,X4,X5)=k2_mcart_1(X5)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, plain, (k11_mcart_1(X1,X2,X3,X4,k4_tarski(k4_tarski(k4_tarski(X5,X6),X7),X8))=X8|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(k4_tarski(k4_tarski(k4_tarski(X5,X6),X7),X8),k4_zfmisc_1(X1,X2,X3,X4))), inference(er,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (X1=k4_mcart_1(esk7_5(X2,X3,X4,X5,X1),esk8_5(X2,X3,X4,X5,X1),esk9_5(X2,X3,X4,X5,X1),esk10_5(X2,X3,X4,X5,X1))|X5=k1_xboole_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|~m1_subset_1(X1,k4_zfmisc_1(X2,X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_19, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6]:(m1_subset_1(X6,k4_zfmisc_1(X1,X2,X3,X4))=>(![X7]:(m1_subset_1(X7,X1)=>![X8]:(m1_subset_1(X8,X2)=>![X9]:(m1_subset_1(X9,X3)=>![X10]:(m1_subset_1(X10,X4)=>(X6=k4_mcart_1(X7,X8,X9,X10)=>X5=X7)))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|(X4=k1_xboole_0|X5=k8_mcart_1(X1,X2,X3,X4,X6)))))))), inference(assume_negation,[status(cth)],[t79_mcart_1])).
cnf(c_0_20, plain, (k2_mcart_1(k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4))=X4|X5=k1_xboole_0|X6=k1_xboole_0|X7=k1_xboole_0|X8=k1_xboole_0|~m1_subset_1(k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4),k4_zfmisc_1(X5,X6,X7,X8))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_21, plain, (X5=k1_xboole_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k4_tarski(k4_tarski(k4_tarski(esk7_5(X2,X3,X4,X5,X1),esk8_5(X2,X3,X4,X5,X1)),esk9_5(X2,X3,X4,X5,X1)),esk10_5(X2,X3,X4,X5,X1))|~m1_subset_1(X1,k4_zfmisc_1(X2,X3,X4,X5))), inference(rw,[status(thm)],[c_0_18, c_0_12])).
fof(c_0_22, negated_conjecture, ![X17, X18, X19, X20]:(m1_subset_1(esk6_0,k4_zfmisc_1(esk1_0,esk2_0,esk3_0,esk4_0))&((~m1_subset_1(X17,esk1_0)|(~m1_subset_1(X18,esk2_0)|(~m1_subset_1(X19,esk3_0)|(~m1_subset_1(X20,esk4_0)|(esk6_0!=k4_mcart_1(X17,X18,X19,X20)|esk5_0=X17)))))&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&(esk4_0!=k1_xboole_0&esk5_0!=k8_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])).
cnf(c_0_23, plain, (esk10_5(X1,X2,X3,X4,X5)=k2_mcart_1(X5)|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|X6=k1_xboole_0|X7=k1_xboole_0|X8=k1_xboole_0|X9=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X9,X8,X7,X6))|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_24, negated_conjecture, (m1_subset_1(esk6_0,k4_zfmisc_1(esk1_0,esk2_0,esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_25, negated_conjecture, (esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, plain, (m1_subset_1(esk10_5(X1,X2,X3,X4,X5),X4)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_30, negated_conjecture, (esk10_5(X1,X2,X3,X4,esk6_0)=k2_mcart_1(esk6_0)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk6_0,k4_zfmisc_1(X1,X2,X3,X4))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28])).
cnf(c_0_31, negated_conjecture, (esk5_0=X1|~m1_subset_1(X1,esk1_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X4,esk4_0)|esk6_0!=k4_mcart_1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32, negated_conjecture, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|m1_subset_1(k2_mcart_1(esk6_0),X4)|~m1_subset_1(esk6_0,k4_zfmisc_1(X1,X2,X3,X4))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_33, negated_conjecture, (esk5_0=X1|esk6_0!=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)|~m1_subset_1(X4,esk4_0)|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk1_0)), inference(rw,[status(thm)],[c_0_31, c_0_12])).
cnf(c_0_34, negated_conjecture, (k4_tarski(k4_tarski(k4_tarski(esk7_5(X1,X2,X3,X4,esk6_0),esk8_5(X1,X2,X3,X4,esk6_0)),esk9_5(X1,X2,X3,X4,esk6_0)),k2_mcart_1(esk6_0))=esk6_0|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|~m1_subset_1(esk6_0,k4_zfmisc_1(X1,X2,X3,X4))), inference(spm,[status(thm)],[c_0_21, c_0_30])).
cnf(c_0_35, negated_conjecture, (m1_subset_1(k2_mcart_1(esk6_0),esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_24]), c_0_28]), c_0_27]), c_0_26]), c_0_25])).
cnf(c_0_36, negated_conjecture, (esk7_5(X1,X2,X3,X4,esk6_0)=esk5_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk9_5(X1,X2,X3,X4,esk6_0),esk3_0)|~m1_subset_1(esk8_5(X1,X2,X3,X4,esk6_0),esk2_0)|~m1_subset_1(esk7_5(X1,X2,X3,X4,esk6_0),esk1_0)|~m1_subset_1(esk6_0,k4_zfmisc_1(X1,X2,X3,X4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])])).
cnf(c_0_37, plain, (m1_subset_1(esk9_5(X1,X2,X3,X4,X5),X3)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_38, plain, (k8_mcart_1(X1,X2,X3,X4,X5)=X6|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X5!=k4_mcart_1(X6,X7,X8,X9)|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_39, negated_conjecture, (esk7_5(X1,X2,esk3_0,X3,esk6_0)=esk5_0|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|~m1_subset_1(esk8_5(X1,X2,esk3_0,X3,esk6_0),esk2_0)|~m1_subset_1(esk7_5(X1,X2,esk3_0,X3,esk6_0),esk1_0)|~m1_subset_1(esk6_0,k4_zfmisc_1(X1,X2,esk3_0,X3))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_26])).
cnf(c_0_40, plain, (m1_subset_1(esk8_5(X1,X2,X3,X4,X5),X2)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_41, plain, (k8_mcart_1(X1,X2,X3,X4,X5)=k1_mcart_1(k1_mcart_1(k1_mcart_1(X5)))|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_42, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k8_mcart_1(X1,X2,X3,X4,X5)=X6|X5!=k4_tarski(k4_tarski(k4_tarski(X6,X7),X8),X9)|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(rw,[status(thm)],[c_0_38, c_0_12])).
cnf(c_0_43, negated_conjecture, (esk7_5(X1,esk2_0,esk3_0,X2,esk6_0)=esk5_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk7_5(X1,esk2_0,esk3_0,X2,esk6_0),esk1_0)|~m1_subset_1(esk6_0,k4_zfmisc_1(X1,esk2_0,esk3_0,X2))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_27]), c_0_26])).
cnf(c_0_44, plain, (m1_subset_1(esk7_5(X1,X2,X3,X4,X5),X1)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_45, negated_conjecture, (esk5_0!=k8_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_46, plain, (k8_mcart_1(X1,X2,X3,X4,X5)=k8_mcart_1(X6,X7,X8,X9,X5)|X6=k1_xboole_0|X7=k1_xboole_0|X8=k1_xboole_0|X9=k1_xboole_0|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))|~m1_subset_1(X5,k4_zfmisc_1(X6,X7,X8,X9))), inference(spm,[status(thm)],[c_0_41, c_0_41])).
cnf(c_0_47, plain, (k8_mcart_1(X1,X2,X3,X4,k4_tarski(k4_tarski(k4_tarski(X5,X6),X7),X8))=X5|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(k4_tarski(k4_tarski(k4_tarski(X5,X6),X7),X8),k4_zfmisc_1(X1,X2,X3,X4))), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_48, negated_conjecture, (esk7_5(esk1_0,esk2_0,esk3_0,X1,esk6_0)=esk5_0|X1=k1_xboole_0|~m1_subset_1(esk6_0,k4_zfmisc_1(esk1_0,esk2_0,esk3_0,X1))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_28]), c_0_26]), c_0_27])).
cnf(c_0_49, negated_conjecture, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|k8_mcart_1(X4,X3,X2,X1,esk6_0)!=esk5_0|~m1_subset_1(esk6_0,k4_zfmisc_1(X4,X3,X2,X1))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_24])]), c_0_28]), c_0_27]), c_0_26]), c_0_25])).
cnf(c_0_50, plain, (k1_mcart_1(k1_mcart_1(k1_mcart_1(k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4))))=X1|X5=k1_xboole_0|X6=k1_xboole_0|X7=k1_xboole_0|X8=k1_xboole_0|~m1_subset_1(k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4),k4_zfmisc_1(X5,X6,X7,X8))), inference(spm,[status(thm)],[c_0_41, c_0_47])).
cnf(c_0_51, negated_conjecture, (k4_tarski(k4_tarski(k4_tarski(esk5_0,esk8_5(esk1_0,esk2_0,esk3_0,X1,esk6_0)),esk9_5(esk1_0,esk2_0,esk3_0,X1,esk6_0)),k2_mcart_1(esk6_0))=esk6_0|X1=k1_xboole_0|~m1_subset_1(esk6_0,k4_zfmisc_1(esk1_0,esk2_0,esk3_0,X1))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_48]), c_0_28]), c_0_27]), c_0_26])).
cnf(c_0_52, negated_conjecture, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|k1_mcart_1(k1_mcart_1(k1_mcart_1(esk6_0)))!=esk5_0|~m1_subset_1(esk6_0,k4_zfmisc_1(X1,X2,X3,X4))), inference(spm,[status(thm)],[c_0_49, c_0_41])).
cnf(c_0_53, negated_conjecture, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|X5=k1_xboole_0|~m1_subset_1(esk6_0,k4_zfmisc_1(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(esk6_0,k4_zfmisc_1(X5,X4,X3,X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52])).
cnf(c_0_54, negated_conjecture, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|~m1_subset_1(esk6_0,k4_zfmisc_1(X1,X2,X3,X4))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_24]), c_0_25])).
cnf(c_0_55, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_24]), c_0_28]), c_0_27]), c_0_26]), c_0_25]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 56
# Proof object clause steps            : 43
# Proof object formula steps           : 13
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 38
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 20
# Processed clauses                    : 77
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 74
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 12
# Backward-rewritten                   : 1
# Generated clauses                    : 170
# ...of the previous two non-trivial   : 168
# Contextual simplify-reflections      : 1
# Paramodulations                      : 165
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 37
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 28
# Current number of unprocessed clauses: 131
# ...number of literals in the above   : 1565
# Current number of archived formulas  : 0
# Current number of archived clauses   : 35
# Clause-clause subsumption calls (NU) : 262
# Rec. Clause-clause subsumption calls : 17
# Non-unit clause-clause subsumptions  : 16
# Unit Clause-clause subsumption calls : 9
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6476

# -------------------------------------------------
# User time                : 0.029 s
# System time              : 0.004 s
# Total time               : 0.033 s
# Maximum resident set size: 28684 pages
