# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t24_yellow_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1)), file('data/MPTP/proofs/1672093db4c74f0ca9805af31d06408f.E_input', t24_yellow_1)).
fof(t56_setfam_1, axiom, ![X1, X2, X3]:((r1_tarski(k3_tarski(X1),X2)&r2_hidden(X3,X1))=>r1_tarski(X3,X2)), file('data/MPTP/proofs/1672093db4c74f0ca9805af31d06408f.E_input', t56_setfam_1)).
fof(t95_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>r1_tarski(k3_tarski(X1),k3_tarski(X2))), file('data/MPTP/proofs/1672093db4c74f0ca9805af31d06408f.E_input', t95_zfmisc_1)).
fof(d1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1))))))))), file('data/MPTP/proofs/1672093db4c74f0ca9805af31d06408f.E_input', d1_pre_topc)).
fof(d12_setfam_1, axiom, ![X1, X2]:(m1_setfam_1(X2,X1)<=>r1_tarski(X1,k3_tarski(X2))), file('data/MPTP/proofs/1672093db4c74f0ca9805af31d06408f.E_input', d12_setfam_1)).
fof(d4_tarski, axiom, ![X1, X2]:(X2=k3_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X3,X4)&r2_hidden(X4,X1)))), file('data/MPTP/proofs/1672093db4c74f0ca9805af31d06408f.E_input', d4_tarski)).
fof(t60_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>(m1_setfam_1(X2,X1)<=>k5_setfam_1(X1,X2)=X1)), file('data/MPTP/proofs/1672093db4c74f0ca9805af31d06408f.E_input', t60_setfam_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/1672093db4c74f0ca9805af31d06408f.E_input', d10_xboole_0)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/1672093db4c74f0ca9805af31d06408f.E_input', d2_subset_1)).
fof(redefinition_k5_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>k5_setfam_1(X1,X2)=k3_tarski(X2)), file('data/MPTP/proofs/1672093db4c74f0ca9805af31d06408f.E_input', redefinition_k5_setfam_1)).
fof(dt_u1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/1672093db4c74f0ca9805af31d06408f.E_input', dt_u1_pre_topc)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/1672093db4c74f0ca9805af31d06408f.E_input', t5_subset)).
fof(t99_zfmisc_1, axiom, ![X1]:k3_tarski(k1_zfmisc_1(X1))=X1, file('data/MPTP/proofs/1672093db4c74f0ca9805af31d06408f.E_input', t99_zfmisc_1)).
fof(t14_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>(r2_hidden(k3_tarski(X1),X1)=>k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1))), file('data/MPTP/proofs/1672093db4c74f0ca9805af31d06408f.E_input', t14_yellow_1)).
fof(c_0_14, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1))), inference(assume_negation,[status(cth)],[t24_yellow_1])).
fof(c_0_15, plain, ![X26, X27, X28]:(~r1_tarski(k3_tarski(X26),X27)|~r2_hidden(X28,X26)|r1_tarski(X28,X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t56_setfam_1])])).
fof(c_0_16, plain, ![X38, X39]:(~r1_tarski(X38,X39)|r1_tarski(k3_tarski(X38),k3_tarski(X39))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t95_zfmisc_1])])).
fof(c_0_17, plain, ![X30, X31, X32, X33]:(((r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|~v2_pre_topc(X30)|~l1_pre_topc(X30))&((~m1_subset_1(X31,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X30))))|(~r1_tarski(X31,u1_pre_topc(X30))|r2_hidden(k5_setfam_1(u1_struct_0(X30),X31),u1_pre_topc(X30)))|~v2_pre_topc(X30)|~l1_pre_topc(X30))&(~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X30)))|(~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X30)))|(~r2_hidden(X32,u1_pre_topc(X30))|~r2_hidden(X33,u1_pre_topc(X30))|r2_hidden(k9_subset_1(u1_struct_0(X30),X32,X33),u1_pre_topc(X30))))|~v2_pre_topc(X30)|~l1_pre_topc(X30))))&(((m1_subset_1(esk6_1(X30),k1_zfmisc_1(u1_struct_0(X30)))|m1_subset_1(esk5_1(X30),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X30))))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30))&((m1_subset_1(esk7_1(X30),k1_zfmisc_1(u1_struct_0(X30)))|m1_subset_1(esk5_1(X30),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X30))))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30))&(((r2_hidden(esk6_1(X30),u1_pre_topc(X30))|m1_subset_1(esk5_1(X30),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X30))))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30))&(r2_hidden(esk7_1(X30),u1_pre_topc(X30))|m1_subset_1(esk5_1(X30),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X30))))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30)))&(~r2_hidden(k9_subset_1(u1_struct_0(X30),esk6_1(X30),esk7_1(X30)),u1_pre_topc(X30))|m1_subset_1(esk5_1(X30),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X30))))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30)))))&(((m1_subset_1(esk6_1(X30),k1_zfmisc_1(u1_struct_0(X30)))|r1_tarski(esk5_1(X30),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30))&((m1_subset_1(esk7_1(X30),k1_zfmisc_1(u1_struct_0(X30)))|r1_tarski(esk5_1(X30),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30))&(((r2_hidden(esk6_1(X30),u1_pre_topc(X30))|r1_tarski(esk5_1(X30),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30))&(r2_hidden(esk7_1(X30),u1_pre_topc(X30))|r1_tarski(esk5_1(X30),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30)))&(~r2_hidden(k9_subset_1(u1_struct_0(X30),esk6_1(X30),esk7_1(X30)),u1_pre_topc(X30))|r1_tarski(esk5_1(X30),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30)))))&((m1_subset_1(esk6_1(X30),k1_zfmisc_1(u1_struct_0(X30)))|~r2_hidden(k5_setfam_1(u1_struct_0(X30),esk5_1(X30)),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30))&((m1_subset_1(esk7_1(X30),k1_zfmisc_1(u1_struct_0(X30)))|~r2_hidden(k5_setfam_1(u1_struct_0(X30),esk5_1(X30)),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30))&(((r2_hidden(esk6_1(X30),u1_pre_topc(X30))|~r2_hidden(k5_setfam_1(u1_struct_0(X30),esk5_1(X30)),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30))&(r2_hidden(esk7_1(X30),u1_pre_topc(X30))|~r2_hidden(k5_setfam_1(u1_struct_0(X30),esk5_1(X30)),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30)))&(~r2_hidden(k9_subset_1(u1_struct_0(X30),esk6_1(X30),esk7_1(X30)),u1_pre_topc(X30))|~r2_hidden(k5_setfam_1(u1_struct_0(X30),esk5_1(X30)),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])).
fof(c_0_18, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&k4_yellow_0(k2_yellow_1(u1_pre_topc(esk1_0)))!=u1_struct_0(esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])).
cnf(c_0_19, plain, (r1_tarski(X3,X2)|~r1_tarski(k3_tarski(X1),X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, plain, (r1_tarski(k3_tarski(X1),k3_tarski(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, plain, (r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_24, plain, ![X9, X10]:((~m1_setfam_1(X10,X9)|r1_tarski(X9,k3_tarski(X10)))&(~r1_tarski(X9,k3_tarski(X10))|m1_setfam_1(X10,X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_setfam_1])])).
cnf(c_0_25, plain, (r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_26, negated_conjecture, (r2_hidden(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])])).
fof(c_0_27, plain, ![X11, X12, X13, X15, X16, X17, X18, X20]:((((r2_hidden(X13,esk2_3(X11,X12,X13))|~r2_hidden(X13,X12)|X12!=k3_tarski(X11))&(r2_hidden(esk2_3(X11,X12,X13),X11)|~r2_hidden(X13,X12)|X12!=k3_tarski(X11)))&(~r2_hidden(X15,X16)|~r2_hidden(X16,X11)|r2_hidden(X15,X12)|X12!=k3_tarski(X11)))&((~r2_hidden(esk3_2(X17,X18),X18)|(~r2_hidden(esk3_2(X17,X18),X20)|~r2_hidden(X20,X17))|X18=k3_tarski(X17))&((r2_hidden(esk3_2(X17,X18),esk4_2(X17,X18))|r2_hidden(esk3_2(X17,X18),X18)|X18=k3_tarski(X17))&(r2_hidden(esk4_2(X17,X18),X17)|r2_hidden(esk3_2(X17,X18),X18)|X18=k3_tarski(X17))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_tarski])])])])])])).
fof(c_0_28, plain, ![X22, X23]:((~m1_setfam_1(X23,X22)|k5_setfam_1(X22,X23)=X22|~m1_subset_1(X23,k1_zfmisc_1(k1_zfmisc_1(X22))))&(k5_setfam_1(X22,X23)!=X22|m1_setfam_1(X23,X22)|~m1_subset_1(X23,k1_zfmisc_1(k1_zfmisc_1(X22))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t60_setfam_1])])])).
cnf(c_0_29, plain, (m1_setfam_1(X2,X1)|~r1_tarski(X1,k3_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, negated_conjecture, (r1_tarski(u1_struct_0(esk1_0),k3_tarski(X1))|~r1_tarski(u1_pre_topc(esk1_0),X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_31, plain, ![X24, X25]:(((r1_tarski(X24,X25)|X24!=X25)&(r1_tarski(X25,X24)|X24!=X25))&(~r1_tarski(X24,X25)|~r1_tarski(X25,X24)|X24=X25)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_32, plain, ![X40, X41]:(((~m1_subset_1(X41,X40)|r2_hidden(X41,X40)|v1_xboole_0(X40))&(~r2_hidden(X41,X40)|m1_subset_1(X41,X40)|v1_xboole_0(X40)))&((~m1_subset_1(X41,X40)|v1_xboole_0(X41)|~v1_xboole_0(X40))&(~v1_xboole_0(X41)|m1_subset_1(X41,X40)|~v1_xboole_0(X40)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
cnf(c_0_33, plain, (r2_hidden(esk2_3(X1,X2,X3),X1)|~r2_hidden(X3,X2)|X2!=k3_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_34, plain, ![X7, X8]:(~m1_subset_1(X8,k1_zfmisc_1(k1_zfmisc_1(X7)))|k5_setfam_1(X7,X8)=k3_tarski(X8)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k5_setfam_1])])).
fof(c_0_35, plain, ![X6]:(~l1_pre_topc(X6)|m1_subset_1(u1_pre_topc(X6),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X6))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_pre_topc])])).
cnf(c_0_36, plain, (k5_setfam_1(X2,X1)=X2|~m1_setfam_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, negated_conjecture, (m1_setfam_1(X1,u1_struct_0(esk1_0))|~r1_tarski(u1_pre_topc(esk1_0),X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_38, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_39, plain, ![X42, X43, X44]:(~r2_hidden(X42,X43)|(~m1_subset_1(X43,k1_zfmisc_1(X44))|~v1_xboole_0(X44))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
cnf(c_0_40, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_41, plain, (r2_hidden(esk2_3(X1,k3_tarski(X1),X2),X1)|~r2_hidden(X2,k3_tarski(X1))), inference(er,[status(thm)],[c_0_33])).
fof(c_0_42, plain, ![X37]:k3_tarski(k1_zfmisc_1(X37))=X37, inference(variable_rename,[status(thm)],[t99_zfmisc_1])).
cnf(c_0_43, plain, (r2_hidden(X1,esk2_3(X2,X3,X1))|~r2_hidden(X1,X3)|X3!=k3_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_44, plain, (k5_setfam_1(X2,X1)=k3_tarski(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_46, negated_conjecture, (k5_setfam_1(u1_struct_0(esk1_0),X1)=u1_struct_0(esk1_0)|~r1_tarski(u1_pre_topc(esk1_0),X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_47, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_48, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_49, plain, (v1_xboole_0(X1)|m1_subset_1(esk2_3(X1,k3_tarski(X1),X2),X1)|~r2_hidden(X2,k3_tarski(X1))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_50, plain, (k3_tarski(k1_zfmisc_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_51, plain, (r2_hidden(X1,esk2_3(X2,k3_tarski(X2),X1))|~r2_hidden(X1,k3_tarski(X2))), inference(er,[status(thm)],[c_0_43])).
fof(c_0_52, plain, ![X29]:(v1_xboole_0(X29)|(~r2_hidden(k3_tarski(X29),X29)|k4_yellow_0(k2_yellow_1(X29))=k3_tarski(X29))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_yellow_1])])])).
cnf(c_0_53, plain, (k5_setfam_1(u1_struct_0(X1),u1_pre_topc(X1))=k3_tarski(u1_pre_topc(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_54, negated_conjecture, (k5_setfam_1(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_45]), c_0_47]), c_0_23])])).
cnf(c_0_55, plain, (v1_xboole_0(k1_zfmisc_1(X1))|~v1_xboole_0(X1)|~r2_hidden(X2,esk2_3(k1_zfmisc_1(X1),X1,X3))|~r2_hidden(X3,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50]), c_0_50])).
cnf(c_0_56, plain, (r2_hidden(X1,esk2_3(k1_zfmisc_1(X2),X2,X1))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_51, c_0_50])).
cnf(c_0_57, plain, (v1_xboole_0(X1)|k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_58, negated_conjecture, (k3_tarski(u1_pre_topc(esk1_0))=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_23])])).
cnf(c_0_59, negated_conjecture, (k4_yellow_0(k2_yellow_1(u1_pre_topc(esk1_0)))!=u1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_60, plain, (v1_xboole_0(k1_zfmisc_1(X1))|~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_61, negated_conjecture, (v1_xboole_0(u1_pre_topc(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_26])]), c_0_59])).
cnf(c_0_62, plain, (m1_subset_1(X1,X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_63, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(u1_pre_topc(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_26]), c_0_61])])).
cnf(c_0_64, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_pre_topc(esk1_0)))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_65, negated_conjecture, (m1_subset_1(u1_pre_topc(esk1_0),k1_zfmisc_1(u1_pre_topc(esk1_0)))), inference(spm,[status(thm)],[c_0_64, c_0_61])).
cnf(c_0_66, negated_conjecture, (~r2_hidden(X1,u1_pre_topc(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_65]), c_0_61])])).
cnf(c_0_67, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_26, c_0_66]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 68
# Proof object clause steps            : 39
# Proof object formula steps           : 29
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 20
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 46
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 46
# Processed clauses                    : 368
# ...of these trivial                  : 2
# ...subsumed                          : 56
# ...remaining for further processing  : 310
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 10
# Backward-rewritten                   : 14
# Generated clauses                    : 765
# ...of the previous two non-trivial   : 666
# Contextual simplify-reflections      : 0
# Paramodulations                      : 759
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 235
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 218
# Current number of unprocessed clauses: 366
# ...number of literals in the above   : 1648
# Current number of archived formulas  : 0
# Current number of archived clauses   : 70
# Clause-clause subsumption calls (NU) : 10189
# Rec. Clause-clause subsumption calls : 1275
# Non-unit clause-clause subsumptions  : 62
# Unit Clause-clause subsumption calls : 259
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 27190

# -------------------------------------------------
# User time                : 0.028 s
# System time              : 0.000 s
# Total time               : 0.028 s
# Maximum resident set size: 28736 pages
