# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t9_yellow_0, conjecture, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(r1_tarski(X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X3,X4)=>r1_lattice3(X1,X2,X4))&(r2_lattice3(X1,X3,X4)=>r2_lattice3(X1,X2,X4)))))), file('data/MPTP/proofs/16ac755a872647a1b3f157afe6b191e4.E_input', t9_yellow_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/16ac755a872647a1b3f157afe6b191e4.E_input', d1_xboole_0)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/16ac755a872647a1b3f157afe6b191e4.E_input', d9_lattice3)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/16ac755a872647a1b3f157afe6b191e4.E_input', t1_xboole_1)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/16ac755a872647a1b3f157afe6b191e4.E_input', d8_lattice3)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/16ac755a872647a1b3f157afe6b191e4.E_input', t5_subset)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/16ac755a872647a1b3f157afe6b191e4.E_input', t3_subset)).
fof(dt_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))), file('data/MPTP/proofs/16ac755a872647a1b3f157afe6b191e4.E_input', dt_k6_domain_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/16ac755a872647a1b3f157afe6b191e4.E_input', d1_tarski)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/16ac755a872647a1b3f157afe6b191e4.E_input', t2_subset)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/16ac755a872647a1b3f157afe6b191e4.E_input', t4_subset)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/16ac755a872647a1b3f157afe6b191e4.E_input', redefinition_k6_domain_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/16ac755a872647a1b3f157afe6b191e4.E_input', d10_xboole_0)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/16ac755a872647a1b3f157afe6b191e4.E_input', t1_subset)).
fof(c_0_14, negated_conjecture, ~(![X1]:(l1_orders_2(X1)=>![X2, X3]:(r1_tarski(X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X3,X4)=>r1_lattice3(X1,X2,X4))&(r2_lattice3(X1,X3,X4)=>r2_lattice3(X1,X2,X4))))))), inference(assume_negation,[status(cth)],[t9_yellow_0])).
fof(c_0_15, plain, ![X9, X10, X11]:((~v1_xboole_0(X9)|~r2_hidden(X10,X9))&(r2_hidden(esk5_1(X11),X11)|v1_xboole_0(X11))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_16, plain, ![X13, X14, X15, X16]:((~r2_lattice3(X13,X14,X15)|(~m1_subset_1(X16,u1_struct_0(X13))|(~r2_hidden(X16,X14)|r1_orders_2(X13,X16,X15)))|~m1_subset_1(X15,u1_struct_0(X13))|~l1_orders_2(X13))&((m1_subset_1(esk6_3(X13,X14,X15),u1_struct_0(X13))|r2_lattice3(X13,X14,X15)|~m1_subset_1(X15,u1_struct_0(X13))|~l1_orders_2(X13))&((r2_hidden(esk6_3(X13,X14,X15),X14)|r2_lattice3(X13,X14,X15)|~m1_subset_1(X15,u1_struct_0(X13))|~l1_orders_2(X13))&(~r1_orders_2(X13,esk6_3(X13,X14,X15),X15)|r2_lattice3(X13,X14,X15)|~m1_subset_1(X15,u1_struct_0(X13))|~l1_orders_2(X13))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
fof(c_0_17, plain, ![X18, X19, X20]:(~r1_tarski(X18,X19)|~r1_tarski(X19,X20)|r1_tarski(X18,X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_18, negated_conjecture, (l1_orders_2(esk1_0)&(r1_tarski(esk2_0,esk3_0)&(m1_subset_1(esk4_0,u1_struct_0(esk1_0))&(((r2_lattice3(esk1_0,esk3_0,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0))&(~r2_lattice3(esk1_0,esk2_0,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)))&((r2_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk1_0,esk2_0,esk4_0))&(~r2_lattice3(esk1_0,esk2_0,esk4_0)|~r1_lattice3(esk1_0,esk2_0,esk4_0))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])])).
cnf(c_0_19, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_21, plain, ![X21, X22, X23, X24]:((~r1_lattice3(X21,X22,X23)|(~m1_subset_1(X24,u1_struct_0(X21))|(~r2_hidden(X24,X22)|r1_orders_2(X21,X23,X24)))|~m1_subset_1(X23,u1_struct_0(X21))|~l1_orders_2(X21))&((m1_subset_1(esk7_3(X21,X22,X23),u1_struct_0(X21))|r1_lattice3(X21,X22,X23)|~m1_subset_1(X23,u1_struct_0(X21))|~l1_orders_2(X21))&((r2_hidden(esk7_3(X21,X22,X23),X22)|r1_lattice3(X21,X22,X23)|~m1_subset_1(X23,u1_struct_0(X21))|~l1_orders_2(X21))&(~r1_orders_2(X21,X23,esk7_3(X21,X22,X23))|r1_lattice3(X21,X22,X23)|~m1_subset_1(X23,u1_struct_0(X21))|~l1_orders_2(X21))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
fof(c_0_22, plain, ![X26, X27, X28]:(~r2_hidden(X26,X27)|(~m1_subset_1(X27,k1_zfmisc_1(X28))|~v1_xboole_0(X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_23, plain, ![X29, X30]:((~m1_subset_1(X29,k1_zfmisc_1(X30))|r1_tarski(X29,X30))&(~r1_tarski(X29,X30)|m1_subset_1(X29,k1_zfmisc_1(X30)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_24, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, negated_conjecture, (r1_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_26, plain, ![X31, X32]:(v1_xboole_0(X31)|~m1_subset_1(X32,X31)|m1_subset_1(k6_domain_1(X31,X32),k1_zfmisc_1(X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_domain_1])])])).
cnf(c_0_27, negated_conjecture, (~r2_lattice3(esk1_0,esk2_0,esk4_0)|~r1_lattice3(esk1_0,esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, plain, (r2_lattice3(X1,X2,X3)|~v1_xboole_0(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_30, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_31, plain, (r2_hidden(esk7_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_32, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_34, plain, ![X33, X34, X35, X36, X37, X38]:(((~r2_hidden(X35,X34)|X35=X33|X34!=k1_tarski(X33))&(X36!=X33|r2_hidden(X36,X34)|X34!=k1_tarski(X33)))&((~r2_hidden(esk8_2(X37,X38),X38)|esk8_2(X37,X38)!=X37|X38=k1_tarski(X37))&(r2_hidden(esk8_2(X37,X38),X38)|esk8_2(X37,X38)=X37|X38=k1_tarski(X37)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_35, negated_conjecture, (r1_tarski(X1,esk3_0)|~r1_tarski(X1,esk2_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_36, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_37, plain, (v1_xboole_0(X1)|m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_38, negated_conjecture, (~v1_xboole_0(esk2_0)|~r1_lattice3(esk1_0,esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30])])).
cnf(c_0_39, plain, (r1_lattice3(X1,X2,X3)|~v1_xboole_0(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_19, c_0_31])).
fof(c_0_40, plain, ![X40, X41]:(~m1_subset_1(X40,X41)|(v1_xboole_0(X41)|r2_hidden(X40,X41))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_41, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X3)|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_42, plain, (r2_hidden(esk8_2(X1,X2),X2)|esk8_2(X1,X2)=X1|X2=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43, plain, (r1_orders_2(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_44, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_45, plain, ![X42, X43, X44]:(~r2_hidden(X42,X43)|~m1_subset_1(X43,k1_zfmisc_1(X44))|m1_subset_1(X42,X44)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_46, negated_conjecture, (r1_tarski(X1,esk3_0)|~r1_tarski(X2,esk2_0)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_24, c_0_35])).
cnf(c_0_47, plain, (v1_xboole_0(X1)|r1_tarski(k6_domain_1(X1,X2),X1)|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_48, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_29]), c_0_30])])).
fof(c_0_49, plain, ![X45, X46]:(v1_xboole_0(X45)|~m1_subset_1(X46,X45)|k6_domain_1(X45,X46)=k1_tarski(X46)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_50, plain, ![X47, X48]:(((r1_tarski(X47,X48)|X47!=X48)&(r1_tarski(X48,X47)|X47!=X48))&(~r1_tarski(X47,X48)|~r1_tarski(X48,X47)|X47=X48)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_51, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_52, plain, (X2=k1_tarski(X1)|~r2_hidden(esk8_2(X1,X2),X2)|esk8_2(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_53, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_54, plain, (esk8_2(X1,X2)=X1|X2=k1_tarski(X1)|~v1_xboole_0(X3)|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
fof(c_0_55, plain, ![X49, X50]:(~r2_hidden(X49,X50)|m1_subset_1(X49,X50)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_56, plain, (r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk6_3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_57, negated_conjecture, (r1_orders_2(esk1_0,X1,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_29]), c_0_30])])).
cnf(c_0_58, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_59, negated_conjecture, (r1_tarski(X1,esk3_0)|~m1_subset_1(X2,esk2_0)|~r1_tarski(X1,k6_domain_1(esk2_0,X2))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])).
cnf(c_0_60, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_61, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_62, plain, (r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_51])])).
cnf(c_0_63, plain, (X1=k1_tarski(X2)|v1_xboole_0(X1)|esk8_2(X2,X1)!=X2|~m1_subset_1(esk8_2(X2,X1),X1)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_64, negated_conjecture, (esk8_2(X1,esk2_0)=X1|k1_tarski(X1)=esk2_0|~v1_xboole_0(esk3_0)), inference(spm,[status(thm)],[c_0_54, c_0_25])).
cnf(c_0_65, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_66, plain, (r2_hidden(esk5_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_67, negated_conjecture, (r2_lattice3(esk1_0,X1,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_hidden(esk6_3(esk1_0,X1,esk4_0),esk3_0)|~m1_subset_1(esk6_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_29]), c_0_30])])).
cnf(c_0_68, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_69, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_58, c_0_33])).
cnf(c_0_70, negated_conjecture, (r1_tarski(X1,esk3_0)|~m1_subset_1(X2,esk2_0)|~r1_tarski(X1,k1_tarski(X2))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_48])).
cnf(c_0_71, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_61])).
cnf(c_0_72, plain, (~v1_xboole_0(X1)|~r1_tarski(k1_tarski(X2),X1)), inference(spm,[status(thm)],[c_0_41, c_0_62])).
cnf(c_0_73, negated_conjecture, (k1_tarski(X1)=esk2_0|~v1_xboole_0(esk3_0)|~m1_subset_1(X1,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_48])).
cnf(c_0_74, plain, (v1_xboole_0(X1)|m1_subset_1(esk5_1(X1),X1)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_75, negated_conjecture, (r2_lattice3(esk1_0,X1,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_hidden(esk6_3(esk1_0,X1,esk4_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_29]), c_0_30])])).
cnf(c_0_76, plain, (m1_subset_1(X1,X2)|~r1_tarski(k1_tarski(X1),X2)), inference(spm,[status(thm)],[c_0_69, c_0_62])).
cnf(c_0_77, negated_conjecture, (r1_tarski(k1_tarski(X1),esk3_0)|~m1_subset_1(X1,esk2_0)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_78, negated_conjecture, (~v1_xboole_0(esk3_0)|~r1_tarski(k1_tarski(X1),esk2_0)), inference(spm,[status(thm)],[c_0_72, c_0_35])).
cnf(c_0_79, negated_conjecture, (k1_tarski(esk5_1(esk2_0))=esk2_0|~v1_xboole_0(esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_48])).
cnf(c_0_80, negated_conjecture, (v1_xboole_0(esk3_0)|r2_lattice3(esk1_0,X1,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)|~m1_subset_1(esk6_3(esk1_0,X1,esk4_0),esk3_0)), inference(spm,[status(thm)],[c_0_75, c_0_53])).
cnf(c_0_81, negated_conjecture, (m1_subset_1(X1,esk3_0)|~m1_subset_1(X1,esk2_0)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_82, negated_conjecture, (~v1_xboole_0(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_71])])).
cnf(c_0_83, negated_conjecture, (r2_lattice3(esk1_0,X1,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)|~m1_subset_1(esk6_3(esk1_0,X1,esk4_0),esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_82])).
cnf(c_0_84, plain, (r2_lattice3(X1,X2,X3)|m1_subset_1(esk6_3(X1,X2,X3),X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_65, c_0_20])).
cnf(c_0_85, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_lattice3(esk1_0,esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_86, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk1_0,esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_87, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_88, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_29]), c_0_30])]), c_0_85])).
cnf(c_0_89, negated_conjecture, (r1_orders_2(esk1_0,X1,esk4_0)|~r2_hidden(X1,esk3_0)|~r1_lattice3(esk1_0,esk2_0,esk4_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_86]), c_0_29]), c_0_30])])).
cnf(c_0_90, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk7_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_91, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,X1)|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_29]), c_0_30])])).
cnf(c_0_92, negated_conjecture, (r2_lattice3(esk1_0,X1,esk4_0)|~r2_hidden(esk6_3(esk1_0,X1,esk4_0),esk3_0)|~r1_lattice3(esk1_0,esk2_0,esk4_0)|~m1_subset_1(esk6_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_89]), c_0_29]), c_0_30])])).
cnf(c_0_93, negated_conjecture, (r1_lattice3(esk1_0,X1,esk4_0)|~r2_hidden(esk7_3(esk1_0,X1,esk4_0),esk3_0)|~m1_subset_1(esk7_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_29]), c_0_30])])).
cnf(c_0_94, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_95, negated_conjecture, (r2_lattice3(esk1_0,X1,esk4_0)|~r2_hidden(esk6_3(esk1_0,X1,esk4_0),esk3_0)|~r1_lattice3(esk1_0,esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_68]), c_0_29]), c_0_30])])).
cnf(c_0_96, negated_conjecture, (r1_lattice3(esk1_0,X1,esk4_0)|~r2_hidden(esk7_3(esk1_0,X1,esk4_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_29]), c_0_30])])).
cnf(c_0_97, negated_conjecture, (v1_xboole_0(esk3_0)|r2_lattice3(esk1_0,X1,esk4_0)|~r1_lattice3(esk1_0,esk2_0,esk4_0)|~m1_subset_1(esk6_3(esk1_0,X1,esk4_0),esk3_0)), inference(spm,[status(thm)],[c_0_95, c_0_53])).
cnf(c_0_98, negated_conjecture, (r1_lattice3(esk1_0,X1,esk4_0)|~m1_subset_1(esk7_3(esk1_0,X1,esk4_0),esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_53]), c_0_82])).
cnf(c_0_99, negated_conjecture, (r2_lattice3(esk1_0,X1,esk4_0)|~r1_lattice3(esk1_0,esk2_0,esk4_0)|~m1_subset_1(esk6_3(esk1_0,X1,esk4_0),esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_81]), c_0_82])).
cnf(c_0_100, negated_conjecture, (r1_lattice3(esk1_0,X1,esk4_0)|~m1_subset_1(esk7_3(esk1_0,X1,esk4_0),esk2_0)), inference(spm,[status(thm)],[c_0_98, c_0_81])).
cnf(c_0_101, plain, (r1_lattice3(X1,X2,X3)|m1_subset_1(esk7_3(X1,X2,X3),X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_65, c_0_31])).
cnf(c_0_102, negated_conjecture, (~r1_lattice3(esk1_0,esk2_0,esk4_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_84]), c_0_29]), c_0_30])]), c_0_27])).
cnf(c_0_103, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_29]), c_0_30])]), c_0_102]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 104
# Proof object clause steps            : 75
# Proof object formula steps           : 29
# Proof object conjectures             : 41
# Proof object clause conjectures      : 38
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 30
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 43
# Proof object simplifying inferences  : 57
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 33
# Processed clauses                    : 153
# ...of these trivial                  : 1
# ...subsumed                          : 44
# ...remaining for further processing  : 108
# Other redundant clauses eliminated   : 9
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 6
# Generated clauses                    : 227
# ...of the previous two non-trivial   : 194
# Contextual simplify-reflections      : 2
# Paramodulations                      : 217
# Factorizations                       : 2
# Equation resolutions                 : 9
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 93
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 81
# Current number of unprocessed clauses: 74
# ...number of literals in the above   : 326
# Current number of archived formulas  : 0
# Current number of archived clauses   : 11
# Clause-clause subsumption calls (NU) : 1515
# Rec. Clause-clause subsumption calls : 847
# Non-unit clause-clause subsumptions  : 28
# Unit Clause-clause subsumption calls : 199
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5384

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.000 s
# Total time               : 0.013 s
# Maximum resident set size: 28740 pages
