# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t31_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>r2_lattice3(X1,k3_waybel_0(X1,X2),X3))))), file('data/MPTP/proofs/16b365478c6e4e37a08986ddd3f75bea.E_input', t31_waybel_0)).
fof(d15_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k3_waybel_0(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>?[X5]:(m1_subset_1(X5,u1_struct_0(X1))&(r1_orders_2(X1,X4,X5)&r2_hidden(X5,X2))))))))), file('data/MPTP/proofs/16b365478c6e4e37a08986ddd3f75bea.E_input', d15_waybel_0)).
fof(dt_k3_waybel_0, axiom, ![X1, X2]:((l1_orders_2(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k3_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/16b365478c6e4e37a08986ddd3f75bea.E_input', dt_k3_waybel_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/16b365478c6e4e37a08986ddd3f75bea.E_input', t24_orders_2)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/16b365478c6e4e37a08986ddd3f75bea.E_input', d9_lattice3)).
fof(t17_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2))))), file('data/MPTP/proofs/16b365478c6e4e37a08986ddd3f75bea.E_input', t17_waybel_0)).
fof(t4_yellow_0, axiom, ![X1]:((v4_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)=>![X4]:((r1_lattice3(X1,X4,X3)=>r1_lattice3(X1,X4,X2))&(r2_lattice3(X1,X4,X2)=>r2_lattice3(X1,X4,X3))))))), file('data/MPTP/proofs/16b365478c6e4e37a08986ddd3f75bea.E_input', t4_yellow_0)).
fof(c_0_7, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>r2_lattice3(X1,k3_waybel_0(X1,X2),X3)))))), inference(assume_negation,[status(cth)],[t31_waybel_0])).
fof(c_0_8, plain, ![X9, X10, X11, X12, X14, X16]:((((m1_subset_1(esk4_4(X9,X10,X11,X12),u1_struct_0(X9))|~r2_hidden(X12,X11)|~m1_subset_1(X12,u1_struct_0(X9))|X11!=k3_waybel_0(X9,X10)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X9)))|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|~l1_orders_2(X9))&((r1_orders_2(X9,X12,esk4_4(X9,X10,X11,X12))|~r2_hidden(X12,X11)|~m1_subset_1(X12,u1_struct_0(X9))|X11!=k3_waybel_0(X9,X10)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X9)))|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|~l1_orders_2(X9))&(r2_hidden(esk4_4(X9,X10,X11,X12),X10)|~r2_hidden(X12,X11)|~m1_subset_1(X12,u1_struct_0(X9))|X11!=k3_waybel_0(X9,X10)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X9)))|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|~l1_orders_2(X9))))&(~m1_subset_1(X14,u1_struct_0(X9))|(~r1_orders_2(X9,X12,X14)|~r2_hidden(X14,X10))|r2_hidden(X12,X11)|~m1_subset_1(X12,u1_struct_0(X9))|X11!=k3_waybel_0(X9,X10)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X9)))|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|~l1_orders_2(X9)))&((m1_subset_1(esk5_3(X9,X10,X11),u1_struct_0(X9))|X11=k3_waybel_0(X9,X10)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X9)))|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|~l1_orders_2(X9))&((~r2_hidden(esk5_3(X9,X10,X11),X11)|(~m1_subset_1(X16,u1_struct_0(X9))|(~r1_orders_2(X9,esk5_3(X9,X10,X11),X16)|~r2_hidden(X16,X10)))|X11=k3_waybel_0(X9,X10)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X9)))|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|~l1_orders_2(X9))&((m1_subset_1(esk6_3(X9,X10,X11),u1_struct_0(X9))|r2_hidden(esk5_3(X9,X10,X11),X11)|X11=k3_waybel_0(X9,X10)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X9)))|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|~l1_orders_2(X9))&((r1_orders_2(X9,esk5_3(X9,X10,X11),esk6_3(X9,X10,X11))|r2_hidden(esk5_3(X9,X10,X11),X11)|X11=k3_waybel_0(X9,X10)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X9)))|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|~l1_orders_2(X9))&(r2_hidden(esk6_3(X9,X10,X11),X10)|r2_hidden(esk5_3(X9,X10,X11),X11)|X11=k3_waybel_0(X9,X10)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X9)))|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|~l1_orders_2(X9))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d15_waybel_0])])])])])).
fof(c_0_9, plain, ![X18, X19]:(~l1_orders_2(X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|m1_subset_1(k3_waybel_0(X18,X19),k1_zfmisc_1(u1_struct_0(X18)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_waybel_0])])).
fof(c_0_10, plain, ![X20, X21]:(v2_struct_0(X20)|(~v3_orders_2(X20)|~l1_orders_2(X20))|(~m1_subset_1(X21,u1_struct_0(X20))|r1_orders_2(X20,X21,X21))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_11, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&((~r2_lattice3(esk1_0,esk2_0,esk3_0)|~r2_lattice3(esk1_0,k3_waybel_0(esk1_0,esk2_0),esk3_0))&(r2_lattice3(esk1_0,esk2_0,esk3_0)|r2_lattice3(esk1_0,k3_waybel_0(esk1_0,esk2_0),esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_7])])])])).
cnf(c_0_12, plain, (r2_hidden(X3,X5)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,X3,X1)|~r2_hidden(X1,X4)|~m1_subset_1(X3,u1_struct_0(X2))|X5!=k3_waybel_0(X2,X4)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (m1_subset_1(k3_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_18, plain, ![X22, X23, X24, X25]:((~r2_lattice3(X22,X23,X24)|(~m1_subset_1(X25,u1_struct_0(X22))|(~r2_hidden(X25,X23)|r1_orders_2(X22,X25,X24)))|~m1_subset_1(X24,u1_struct_0(X22))|~l1_orders_2(X22))&((m1_subset_1(esk7_3(X22,X23,X24),u1_struct_0(X22))|r2_lattice3(X22,X23,X24)|~m1_subset_1(X24,u1_struct_0(X22))|~l1_orders_2(X22))&((r2_hidden(esk7_3(X22,X23,X24),X23)|r2_lattice3(X22,X23,X24)|~m1_subset_1(X24,u1_struct_0(X22))|~l1_orders_2(X22))&(~r1_orders_2(X22,esk7_3(X22,X23,X24),X24)|r2_lattice3(X22,X23,X24)|~m1_subset_1(X24,u1_struct_0(X22))|~l1_orders_2(X22))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
cnf(c_0_19, plain, (r2_hidden(X1,k3_waybel_0(X2,X3))|k3_waybel_0(X2,X3)!=k3_waybel_0(X2,X4)|~r1_orders_2(X2,X1,X5)|~r2_hidden(X5,X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_20, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])]), c_0_17])).
cnf(c_0_21, plain, (r1_orders_2(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk3_0)|r2_lattice3(esk1_0,k3_waybel_0(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_23, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_24, negated_conjecture, (r2_hidden(X1,k3_waybel_0(esk1_0,X2))|k3_waybel_0(esk1_0,X2)!=k3_waybel_0(esk1_0,X3)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_16])])).
cnf(c_0_25, plain, (r2_hidden(esk7_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, negated_conjecture, (r1_orders_2(esk1_0,X1,esk3_0)|r2_lattice3(esk1_0,esk2_0,esk3_0)|~r2_hidden(X1,k3_waybel_0(esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_16])])).
cnf(c_0_27, negated_conjecture, (r2_hidden(esk7_3(X1,X2,X3),k3_waybel_0(esk1_0,X4))|r2_lattice3(X1,X2,X3)|k3_waybel_0(esk1_0,X4)!=k3_waybel_0(esk1_0,X2)|~m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_28, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_29, plain, (r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk7_3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_30, negated_conjecture, (r1_orders_2(esk1_0,esk7_3(X1,X2,X3),esk3_0)|r2_lattice3(esk1_0,esk2_0,esk3_0)|r2_lattice3(X1,X2,X3)|k3_waybel_0(esk1_0,esk2_0)!=k3_waybel_0(esk1_0,X2)|~m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])])).
cnf(c_0_31, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk3_0)|r2_lattice3(esk1_0,X1,esk3_0)|k3_waybel_0(esk1_0,esk2_0)!=k3_waybel_0(esk1_0,X1)|~m1_subset_1(esk7_3(esk1_0,X1,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_23]), c_0_16])])).
cnf(c_0_32, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_33, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk3_0)|r2_lattice3(esk1_0,X1,esk3_0)|k3_waybel_0(esk1_0,esk2_0)!=k3_waybel_0(esk1_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_23]), c_0_16])])).
cnf(c_0_34, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_33]), c_0_28])])).
fof(c_0_35, plain, ![X27, X28, X29]:((~r2_hidden(X29,k5_waybel_0(X27,X28))|r1_orders_2(X27,X29,X28)|~m1_subset_1(X29,u1_struct_0(X27))|~m1_subset_1(X28,u1_struct_0(X27))|(v2_struct_0(X27)|~l1_orders_2(X27)))&(~r1_orders_2(X27,X29,X28)|r2_hidden(X29,k5_waybel_0(X27,X28))|~m1_subset_1(X29,u1_struct_0(X27))|~m1_subset_1(X28,u1_struct_0(X27))|(v2_struct_0(X27)|~l1_orders_2(X27)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_waybel_0])])])])])).
fof(c_0_36, plain, ![X30, X31, X32, X33]:((~r1_lattice3(X30,X33,X32)|r1_lattice3(X30,X33,X31)|~r1_orders_2(X30,X31,X32)|~m1_subset_1(X32,u1_struct_0(X30))|~m1_subset_1(X31,u1_struct_0(X30))|(~v4_orders_2(X30)|~l1_orders_2(X30)))&(~r2_lattice3(X30,X33,X31)|r2_lattice3(X30,X33,X32)|~r1_orders_2(X30,X31,X32)|~m1_subset_1(X32,u1_struct_0(X30))|~m1_subset_1(X31,u1_struct_0(X30))|(~v4_orders_2(X30)|~l1_orders_2(X30)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_yellow_0])])])])).
cnf(c_0_37, negated_conjecture, (r1_orders_2(esk1_0,X1,esk3_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_34]), c_0_23]), c_0_16])])).
cnf(c_0_38, plain, (r2_hidden(esk4_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_39, plain, (r1_orders_2(X2,X1,X3)|v2_struct_0(X2)|~r2_hidden(X1,k5_waybel_0(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_40, plain, (r2_lattice3(X1,X2,X4)|~r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,X3,X4)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_41, negated_conjecture, (r1_orders_2(esk1_0,esk4_4(X1,esk2_0,X2,X3),esk3_0)|X2!=k3_waybel_0(X1,esk2_0)|~r2_hidden(X3,X2)|~m1_subset_1(esk4_4(X1,esk2_0,X2,X3),u1_struct_0(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_42, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_43, plain, (r1_orders_2(X1,esk7_3(X2,k5_waybel_0(X1,X3),X4),X3)|r2_lattice3(X2,k5_waybel_0(X1,X3),X4)|v2_struct_0(X1)|~m1_subset_1(esk7_3(X2,k5_waybel_0(X1,X3),X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X1)|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_39, c_0_25])).
cnf(c_0_44, negated_conjecture, (r2_lattice3(esk1_0,X1,esk3_0)|X2!=k3_waybel_0(X3,esk2_0)|~r2_hidden(X4,X2)|~r2_lattice3(esk1_0,X1,esk4_4(X3,esk2_0,X2,X4))|~m1_subset_1(esk4_4(X3,esk2_0,X2,X4),u1_struct_0(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_orders_2(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_23]), c_0_16]), c_0_42])])).
cnf(c_0_45, plain, (r2_lattice3(X1,k5_waybel_0(X1,X2),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_43]), c_0_32])).
cnf(c_0_46, negated_conjecture, (r2_lattice3(esk1_0,k5_waybel_0(esk1_0,esk4_4(X1,esk2_0,X2,X3)),esk3_0)|X2!=k3_waybel_0(X1,esk2_0)|~r2_hidden(X3,X2)|~m1_subset_1(esk4_4(X1,esk2_0,X2,X3),u1_struct_0(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_16])]), c_0_17])).
cnf(c_0_47, plain, (r2_hidden(X2,k5_waybel_0(X1,X3))|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_48, plain, (r1_orders_2(X1,X2,esk4_4(X1,X3,X4,X2))|~r2_hidden(X2,X4)|~m1_subset_1(X2,u1_struct_0(X1))|X4!=k3_waybel_0(X1,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_49, plain, (m1_subset_1(esk4_4(X1,X2,X3,X4),u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_50, negated_conjecture, (r1_orders_2(esk1_0,X1,esk3_0)|X2!=k3_waybel_0(X3,esk2_0)|~r2_hidden(X1,k5_waybel_0(esk1_0,esk4_4(X3,esk2_0,X2,X4)))|~r2_hidden(X4,X2)|~m1_subset_1(esk4_4(X3,esk2_0,X2,X4),u1_struct_0(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_orders_2(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_46]), c_0_23]), c_0_16])])).
cnf(c_0_51, plain, (r2_hidden(X1,k5_waybel_0(X2,esk4_4(X2,X3,X4,X1)))|v2_struct_0(X2)|X4!=k3_waybel_0(X2,X3)|~r2_hidden(X1,X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])).
cnf(c_0_52, negated_conjecture, (r1_orders_2(esk1_0,X1,esk3_0)|X2!=k3_waybel_0(esk1_0,esk2_0)|~r2_hidden(X1,X2)|~m1_subset_1(esk4_4(esk1_0,esk2_0,X2,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_28]), c_0_16])]), c_0_17])).
cnf(c_0_53, negated_conjecture, (r1_orders_2(esk1_0,X1,esk3_0)|X2!=k3_waybel_0(esk1_0,esk2_0)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_49]), c_0_28]), c_0_16])])).
cnf(c_0_54, negated_conjecture, (r1_orders_2(esk1_0,X1,esk3_0)|k3_waybel_0(esk1_0,X2)!=k3_waybel_0(esk1_0,esk2_0)|~r2_hidden(X1,k3_waybel_0(esk1_0,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_13]), c_0_16])])).
cnf(c_0_55, negated_conjecture, (r1_orders_2(esk1_0,esk7_3(X1,k3_waybel_0(esk1_0,X2),X3),esk3_0)|r2_lattice3(X1,k3_waybel_0(esk1_0,X2),X3)|k3_waybel_0(esk1_0,X2)!=k3_waybel_0(esk1_0,esk2_0)|~m1_subset_1(esk7_3(X1,k3_waybel_0(esk1_0,X2),X3),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_54, c_0_25])).
cnf(c_0_56, negated_conjecture, (r2_lattice3(esk1_0,k3_waybel_0(esk1_0,X1),esk3_0)|k3_waybel_0(esk1_0,X1)!=k3_waybel_0(esk1_0,esk2_0)|~m1_subset_1(esk7_3(esk1_0,k3_waybel_0(esk1_0,X1),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_55]), c_0_23]), c_0_16])])).
cnf(c_0_57, negated_conjecture, (~r2_lattice3(esk1_0,esk2_0,esk3_0)|~r2_lattice3(esk1_0,k3_waybel_0(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_58, negated_conjecture, (r2_lattice3(esk1_0,k3_waybel_0(esk1_0,X1),esk3_0)|k3_waybel_0(esk1_0,X1)!=k3_waybel_0(esk1_0,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_32]), c_0_23]), c_0_16])])).
cnf(c_0_59, negated_conjecture, (~r2_lattice3(esk1_0,k3_waybel_0(esk1_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_34])])).
cnf(c_0_60, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_58]), c_0_28])]), c_0_59]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 61
# Proof object clause steps            : 46
# Proof object formula steps           : 15
# Proof object conjectures             : 32
# Proof object clause conjectures      : 29
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 24
# Proof object simplifying inferences  : 53
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 251
# ...of these trivial                  : 1
# ...subsumed                          : 46
# ...remaining for further processing  : 204
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 9
# Backward-rewritten                   : 11
# Generated clauses                    : 348
# ...of the previous two non-trivial   : 338
# Contextual simplify-reflections      : 15
# Paramodulations                      : 334
# Factorizations                       : 0
# Equation resolutions                 : 14
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 157
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 149
# Current number of unprocessed clauses: 135
# ...number of literals in the above   : 1395
# Current number of archived formulas  : 0
# Current number of archived clauses   : 47
# Clause-clause subsumption calls (NU) : 6677
# Rec. Clause-clause subsumption calls : 410
# Non-unit clause-clause subsumptions  : 70
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 22493

# -------------------------------------------------
# User time                : 0.027 s
# System time              : 0.000 s
# Total time               : 0.027 s
# Maximum resident set size: 28684 pages
