# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t53_lattices, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v17_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_lattices(X1,X2,X3)=>r3_lattices(X1,k7_lattices(X1,X3),k7_lattices(X1,X2)))))), file('data/MPTP/proofs/16bcb0e5cc4a4e5586511955a89de5ca.E_input', t53_lattices)).
fof(t49_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v17_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>k7_lattices(X1,k7_lattices(X1,X2))=X2)), file('data/MPTP/proofs/16bcb0e5cc4a4e5586511955a89de5ca.E_input', t49_lattices)).
fof(t51_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v17_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k7_lattices(X1,k3_lattices(X1,X2,X3))=k4_lattices(X1,k7_lattices(X1,X2),k7_lattices(X1,X3))))), file('data/MPTP/proofs/16bcb0e5cc4a4e5586511955a89de5ca.E_input', t51_lattices)).
fof(t52_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v17_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k4_lattices(X1,X2,X3)=k5_lattices(X1)<=>r3_lattices(X1,X2,k7_lattices(X1,X3)))))), file('data/MPTP/proofs/16bcb0e5cc4a4e5586511955a89de5ca.E_input', t52_lattices)).
fof(commutativity_k3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v4_lattices(X1)&l2_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k3_lattices(X1,X2,X3)=k3_lattices(X1,X3,X2)), file('data/MPTP/proofs/16bcb0e5cc4a4e5586511955a89de5ca.E_input', commutativity_k3_lattices)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/16bcb0e5cc4a4e5586511955a89de5ca.E_input', cc1_lattices)).
fof(dt_l3_lattices, axiom, ![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1))), file('data/MPTP/proofs/16bcb0e5cc4a4e5586511955a89de5ca.E_input', dt_l3_lattices)).
fof(dt_k7_lattices, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l3_lattices(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k7_lattices(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/16bcb0e5cc4a4e5586511955a89de5ca.E_input', dt_k7_lattices)).
fof(c_0_8, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v17_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_lattices(X1,X2,X3)=>r3_lattices(X1,k7_lattices(X1,X3),k7_lattices(X1,X2))))))), inference(assume_negation,[status(cth)],[t53_lattices])).
fof(c_0_9, plain, ![X7, X8]:(v2_struct_0(X7)|(~v10_lattices(X7)|(~v17_lattices(X7)|~l3_lattices(X7)))|(~m1_subset_1(X8,u1_struct_0(X7))|k7_lattices(X7,k7_lattices(X7,X8))=X8)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t49_lattices])])])])).
fof(c_0_10, negated_conjecture, ((~v2_struct_0(esk1_0)&(v10_lattices(esk1_0)&(v17_lattices(esk1_0)&l3_lattices(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(r3_lattices(esk1_0,esk2_0,esk3_0)&~r3_lattices(esk1_0,k7_lattices(esk1_0,esk3_0),k7_lattices(esk1_0,esk2_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])).
fof(c_0_11, plain, ![X9, X10, X11]:(v2_struct_0(X9)|(~v10_lattices(X9)|(~v17_lattices(X9)|~l3_lattices(X9)))|(~m1_subset_1(X10,u1_struct_0(X9))|(~m1_subset_1(X11,u1_struct_0(X9))|k7_lattices(X9,k3_lattices(X9,X10,X11))=k4_lattices(X9,k7_lattices(X9,X10),k7_lattices(X9,X11))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t51_lattices])])])])).
cnf(c_0_12, plain, (v2_struct_0(X1)|k7_lattices(X1,k7_lattices(X1,X2))=X2|~v10_lattices(X1)|~v17_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, negated_conjecture, (v17_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, negated_conjecture, (l3_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (v10_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, plain, (v2_struct_0(X1)|k7_lattices(X1,k3_lattices(X1,X2,X3))=k4_lattices(X1,k7_lattices(X1,X2),k7_lattices(X1,X3))|~v10_lattices(X1)|~v17_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, negated_conjecture, (k7_lattices(esk1_0,k7_lattices(esk1_0,X1))=X1|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_15])]), c_0_16])).
fof(c_0_19, plain, ![X12, X13, X14]:((k4_lattices(X12,X13,X14)!=k5_lattices(X12)|r3_lattices(X12,X13,k7_lattices(X12,X14))|~m1_subset_1(X14,u1_struct_0(X12))|~m1_subset_1(X13,u1_struct_0(X12))|(v2_struct_0(X12)|(~v10_lattices(X12)|(~v17_lattices(X12)|~l3_lattices(X12)))))&(~r3_lattices(X12,X13,k7_lattices(X12,X14))|k4_lattices(X12,X13,X14)=k5_lattices(X12)|~m1_subset_1(X14,u1_struct_0(X12))|~m1_subset_1(X13,u1_struct_0(X12))|(v2_struct_0(X12)|(~v10_lattices(X12)|(~v17_lattices(X12)|~l3_lattices(X12)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t52_lattices])])])])])).
cnf(c_0_20, negated_conjecture, (k4_lattices(esk1_0,k7_lattices(esk1_0,X1),X2)=k7_lattices(esk1_0,k3_lattices(esk1_0,X1,k7_lattices(esk1_0,X2)))|~m1_subset_1(k7_lattices(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_14]), c_0_13]), c_0_15])]), c_0_16])).
cnf(c_0_21, plain, (r3_lattices(X1,X2,k7_lattices(X1,X3))|v2_struct_0(X1)|k4_lattices(X1,X2,X3)!=k5_lattices(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v17_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_22, negated_conjecture, (k4_lattices(esk1_0,X1,X2)=k7_lattices(esk1_0,k3_lattices(esk1_0,k7_lattices(esk1_0,X1),k7_lattices(esk1_0,X2)))|~m1_subset_1(k7_lattices(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(k7_lattices(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_20, c_0_18])).
fof(c_0_23, plain, ![X15, X16, X17]:(v2_struct_0(X15)|(~v4_lattices(X15)|~l2_lattices(X15))|(~m1_subset_1(X16,u1_struct_0(X15))|~m1_subset_1(X17,u1_struct_0(X15)))|k3_lattices(X15,X16,X17)=k3_lattices(X15,X17,X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[commutativity_k3_lattices])])])).
fof(c_0_24, plain, ![X18]:((~v2_struct_0(X18)|(v2_struct_0(X18)|~v10_lattices(X18))|~l3_lattices(X18))&((v4_lattices(X18)|(v2_struct_0(X18)|~v10_lattices(X18))|~l3_lattices(X18))&((v5_lattices(X18)|(v2_struct_0(X18)|~v10_lattices(X18))|~l3_lattices(X18))&((v6_lattices(X18)|(v2_struct_0(X18)|~v10_lattices(X18))|~l3_lattices(X18))&((v7_lattices(X18)|(v2_struct_0(X18)|~v10_lattices(X18))|~l3_lattices(X18))&((v8_lattices(X18)|(v2_struct_0(X18)|~v10_lattices(X18))|~l3_lattices(X18))&(v9_lattices(X18)|(v2_struct_0(X18)|~v10_lattices(X18))|~l3_lattices(X18)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
fof(c_0_25, plain, ![X19]:((l1_lattices(X19)|~l3_lattices(X19))&(l2_lattices(X19)|~l3_lattices(X19))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
cnf(c_0_26, negated_conjecture, (r3_lattices(esk1_0,X1,k7_lattices(esk1_0,X2))|k7_lattices(esk1_0,k3_lattices(esk1_0,k7_lattices(esk1_0,X1),k7_lattices(esk1_0,X2)))!=k5_lattices(esk1_0)|~m1_subset_1(k7_lattices(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(k7_lattices(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_14]), c_0_13]), c_0_15])]), c_0_16])).
cnf(c_0_27, plain, (v2_struct_0(X1)|k3_lattices(X1,X2,X3)=k3_lattices(X1,X3,X2)|~v4_lattices(X1)|~l2_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28, plain, (v4_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29, plain, (l2_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30, negated_conjecture, (r3_lattices(esk1_0,X1,X2)|k7_lattices(esk1_0,k3_lattices(esk1_0,k7_lattices(esk1_0,X1),X2))!=k5_lattices(esk1_0)|~m1_subset_1(k7_lattices(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(k7_lattices(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_26, c_0_18])).
cnf(c_0_31, plain, (k3_lattices(X1,X2,X3)=k3_lattices(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])).
cnf(c_0_32, plain, (k4_lattices(X1,X2,X3)=k5_lattices(X1)|v2_struct_0(X1)|~r3_lattices(X1,X2,k7_lattices(X1,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v17_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_33, negated_conjecture, (r3_lattices(esk1_0,X1,X2)|k7_lattices(esk1_0,k3_lattices(esk1_0,X2,k7_lattices(esk1_0,X1)))!=k5_lattices(esk1_0)|~m1_subset_1(k7_lattices(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(k7_lattices(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_34, negated_conjecture, (k7_lattices(esk1_0,k3_lattices(esk1_0,X1,k7_lattices(esk1_0,X2)))=k5_lattices(esk1_0)|~r3_lattices(esk1_0,k7_lattices(esk1_0,X1),k7_lattices(esk1_0,X2))|~m1_subset_1(k7_lattices(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(k7_lattices(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_20]), c_0_14]), c_0_13]), c_0_15])]), c_0_16])).
cnf(c_0_35, negated_conjecture, (r3_lattices(esk1_0,X1,X2)|~r3_lattices(esk1_0,k7_lattices(esk1_0,X2),k7_lattices(esk1_0,X1))|~m1_subset_1(k7_lattices(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(k7_lattices(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_36, negated_conjecture, (r3_lattices(esk1_0,k7_lattices(esk1_0,X1),X2)|~r3_lattices(esk1_0,k7_lattices(esk1_0,X2),X1)|~m1_subset_1(k7_lattices(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(k7_lattices(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_35, c_0_18])).
cnf(c_0_37, negated_conjecture, (r3_lattices(esk1_0,k7_lattices(esk1_0,X1),k7_lattices(esk1_0,X2))|~r3_lattices(esk1_0,X2,X1)|~m1_subset_1(k7_lattices(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(k7_lattices(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_36, c_0_18])).
cnf(c_0_38, negated_conjecture, (r3_lattices(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_39, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_40, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_41, negated_conjecture, (~r3_lattices(esk1_0,k7_lattices(esk1_0,esk3_0),k7_lattices(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_42, plain, ![X20, X21]:(v2_struct_0(X20)|~l3_lattices(X20)|~m1_subset_1(X21,u1_struct_0(X20))|m1_subset_1(k7_lattices(X20,X21),u1_struct_0(X20))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_lattices])])])).
cnf(c_0_43, negated_conjecture, (~m1_subset_1(k7_lattices(esk1_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k7_lattices(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39]), c_0_40])]), c_0_41])).
cnf(c_0_44, plain, (v2_struct_0(X1)|m1_subset_1(k7_lattices(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_45, negated_conjecture, (~m1_subset_1(k7_lattices(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_40]), c_0_14])]), c_0_16])).
cnf(c_0_46, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_44]), c_0_39]), c_0_14])]), c_0_16]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 47
# Proof object clause steps            : 30
# Proof object formula steps           : 17
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 36
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 22
# Processed clauses                    : 113
# ...of these trivial                  : 0
# ...subsumed                          : 29
# ...remaining for further processing  : 84
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 363
# ...of the previous two non-trivial   : 344
# Contextual simplify-reflections      : 5
# Paramodulations                      : 362
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 62
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 53
# Current number of unprocessed clauses: 275
# ...number of literals in the above   : 2425
# Current number of archived formulas  : 0
# Current number of archived clauses   : 22
# Clause-clause subsumption calls (NU) : 1122
# Rec. Clause-clause subsumption calls : 193
# Non-unit clause-clause subsumptions  : 34
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 19490

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 28700 pages
