# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t39_pre_topc, conjecture, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/16e3733294184c72ba665f50b463ce7a.E_input', t39_pre_topc)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/16e3733294184c72ba665f50b463ce7a.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/16e3733294184c72ba665f50b463ce7a.E_input', d4_subset_1)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/16e3733294184c72ba665f50b463ce7a.E_input', d1_zfmisc_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/16e3733294184c72ba665f50b463ce7a.E_input', commutativity_k2_xboole_0)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/16e3733294184c72ba665f50b463ce7a.E_input', t12_xboole_1)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/16e3733294184c72ba665f50b463ce7a.E_input', d3_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/16e3733294184c72ba665f50b463ce7a.E_input', dt_l1_pre_topc)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/16e3733294184c72ba665f50b463ce7a.E_input', dt_m1_pre_topc)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/16e3733294184c72ba665f50b463ce7a.E_input', t5_subset)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/16e3733294184c72ba665f50b463ce7a.E_input', t36_xboole_1)).
fof(t10_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(X1,k2_xboole_0(X3,X2))), file('data/MPTP/proofs/16e3733294184c72ba665f50b463ce7a.E_input', t10_xboole_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/16e3733294184c72ba665f50b463ce7a.E_input', t2_subset)).
fof(d9_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>(m1_pre_topc(X2,X1)<=>(r1_tarski(k2_struct_0(X2),k2_struct_0(X1))&![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>(r2_hidden(X3,u1_pre_topc(X2))<=>?[X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))&(r2_hidden(X4,u1_pre_topc(X1))&X3=k9_subset_1(u1_struct_0(X2),X4,k2_struct_0(X2)))))))))), file('data/MPTP/proofs/16e3733294184c72ba665f50b463ce7a.E_input', d9_pre_topc)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/16e3733294184c72ba665f50b463ce7a.E_input', d2_subset_1)).
fof(c_0_15, negated_conjecture, ~(![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))))), inference(assume_negation,[status(cth)],[t39_pre_topc])).
fof(c_0_16, plain, ![X8]:m1_subset_1(k2_subset_1(X8),k1_zfmisc_1(X8)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_17, plain, ![X9]:k2_subset_1(X9)=X9, inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_18, plain, ![X10, X11, X12, X13, X14, X15]:(((~r2_hidden(X12,X11)|r1_tarski(X12,X10)|X11!=k1_zfmisc_1(X10))&(~r1_tarski(X13,X10)|r2_hidden(X13,X11)|X11!=k1_zfmisc_1(X10)))&((~r2_hidden(esk4_2(X14,X15),X15)|~r1_tarski(esk4_2(X14,X15),X14)|X15=k1_zfmisc_1(X14))&(r2_hidden(esk4_2(X14,X15),X15)|r1_tarski(esk4_2(X14,X15),X14)|X15=k1_zfmisc_1(X14)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
fof(c_0_19, plain, ![X17, X18]:k2_xboole_0(X17,X18)=k2_xboole_0(X18,X17), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_20, plain, ![X19, X20]:(~r1_tarski(X19,X20)|k2_xboole_0(X19,X20)=X20), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
fof(c_0_21, plain, ![X21]:(~l1_struct_0(X21)|k2_struct_0(X21)=u1_struct_0(X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
fof(c_0_22, plain, ![X22]:(~l1_pre_topc(X22)|l1_struct_0(X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_23, plain, ![X23, X24]:(~l1_pre_topc(X23)|(~m1_pre_topc(X24,X23)|l1_pre_topc(X24))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
fof(c_0_24, negated_conjecture, (l1_pre_topc(esk1_0)&(m1_pre_topc(esk2_0,esk1_0)&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).
fof(c_0_25, plain, ![X25, X26, X27]:(~r2_hidden(X25,X26)|(~m1_subset_1(X26,k1_zfmisc_1(X27))|~v1_xboole_0(X27))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
cnf(c_0_26, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_27, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28, plain, (r2_hidden(X1,X3)|~r1_tarski(X1,X2)|X3!=k1_zfmisc_1(X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_29, plain, ![X28, X29]:r1_tarski(k4_xboole_0(X28,X29),X28), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_30, plain, ![X30, X31, X32]:(~r1_tarski(X30,X31)|r1_tarski(X30,k2_xboole_0(X32,X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_xboole_1])])).
cnf(c_0_31, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_32, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_33, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_34, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_35, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_36, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_37, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_38, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_39, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_40, plain, (r2_hidden(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(er,[status(thm)],[c_0_28])).
cnf(c_0_41, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_42, plain, (r1_tarski(X1,k2_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_43, plain, (k2_xboole_0(X1,X2)=X1|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_44, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_45, plain, ![X33, X34]:(~m1_subset_1(X33,X34)|(v1_xboole_0(X34)|r2_hidden(X33,X34))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_47, plain, (u1_struct_0(X1)=k2_struct_0(X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_48, negated_conjecture, (l1_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37])])).
cnf(c_0_49, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_50, plain, (r2_hidden(k4_xboole_0(X1,X2),k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_51, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,X3)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_52, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_53, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])])).
cnf(c_0_55, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
fof(c_0_56, plain, ![X35, X36, X37, X39, X41]:(((r1_tarski(k2_struct_0(X36),k2_struct_0(X35))|~m1_pre_topc(X36,X35)|~l1_pre_topc(X36)|~l1_pre_topc(X35))&(((m1_subset_1(esk5_3(X35,X36,X37),k1_zfmisc_1(u1_struct_0(X35)))|~r2_hidden(X37,u1_pre_topc(X36))|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X36)))|~m1_pre_topc(X36,X35)|~l1_pre_topc(X36)|~l1_pre_topc(X35))&((r2_hidden(esk5_3(X35,X36,X37),u1_pre_topc(X35))|~r2_hidden(X37,u1_pre_topc(X36))|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X36)))|~m1_pre_topc(X36,X35)|~l1_pre_topc(X36)|~l1_pre_topc(X35))&(X37=k9_subset_1(u1_struct_0(X36),esk5_3(X35,X36,X37),k2_struct_0(X36))|~r2_hidden(X37,u1_pre_topc(X36))|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X36)))|~m1_pre_topc(X36,X35)|~l1_pre_topc(X36)|~l1_pre_topc(X35))))&(~m1_subset_1(X39,k1_zfmisc_1(u1_struct_0(X35)))|(~r2_hidden(X39,u1_pre_topc(X35))|X37!=k9_subset_1(u1_struct_0(X36),X39,k2_struct_0(X36)))|r2_hidden(X37,u1_pre_topc(X36))|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X36)))|~m1_pre_topc(X36,X35)|~l1_pre_topc(X36)|~l1_pre_topc(X35))))&((m1_subset_1(esk6_2(X35,X36),k1_zfmisc_1(u1_struct_0(X36)))|~r1_tarski(k2_struct_0(X36),k2_struct_0(X35))|m1_pre_topc(X36,X35)|~l1_pre_topc(X36)|~l1_pre_topc(X35))&((~r2_hidden(esk6_2(X35,X36),u1_pre_topc(X36))|(~m1_subset_1(X41,k1_zfmisc_1(u1_struct_0(X35)))|(~r2_hidden(X41,u1_pre_topc(X35))|esk6_2(X35,X36)!=k9_subset_1(u1_struct_0(X36),X41,k2_struct_0(X36))))|~r1_tarski(k2_struct_0(X36),k2_struct_0(X35))|m1_pre_topc(X36,X35)|~l1_pre_topc(X36)|~l1_pre_topc(X35))&((m1_subset_1(esk7_2(X35,X36),k1_zfmisc_1(u1_struct_0(X35)))|r2_hidden(esk6_2(X35,X36),u1_pre_topc(X36))|~r1_tarski(k2_struct_0(X36),k2_struct_0(X35))|m1_pre_topc(X36,X35)|~l1_pre_topc(X36)|~l1_pre_topc(X35))&((r2_hidden(esk7_2(X35,X36),u1_pre_topc(X35))|r2_hidden(esk6_2(X35,X36),u1_pre_topc(X36))|~r1_tarski(k2_struct_0(X36),k2_struct_0(X35))|m1_pre_topc(X36,X35)|~l1_pre_topc(X36)|~l1_pre_topc(X35))&(esk6_2(X35,X36)=k9_subset_1(u1_struct_0(X36),esk7_2(X35,X36),k2_struct_0(X36))|r2_hidden(esk6_2(X35,X36),u1_pre_topc(X36))|~r1_tarski(k2_struct_0(X36),k2_struct_0(X35))|m1_pre_topc(X36,X35)|~l1_pre_topc(X36)|~l1_pre_topc(X35))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_pre_topc])])])])])).
fof(c_0_57, plain, ![X43, X44]:(((~m1_subset_1(X44,X43)|r2_hidden(X44,X43)|v1_xboole_0(X43))&(~r2_hidden(X44,X43)|m1_subset_1(X44,X43)|v1_xboole_0(X43)))&((~m1_subset_1(X44,X43)|v1_xboole_0(X44)|~v1_xboole_0(X43))&(~v1_xboole_0(X44)|m1_subset_1(X44,X43)|~v1_xboole_0(X43)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
cnf(c_0_58, plain, (r1_tarski(X1,X2)|~r1_tarski(X3,X2)|~r2_hidden(X1,k1_zfmisc_1(X3))), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_59, negated_conjecture, (r2_hidden(esk3_0,k1_zfmisc_1(k2_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55])).
cnf(c_0_60, plain, (r1_tarski(k2_struct_0(X1),k2_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X1)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_61, negated_conjecture, (~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_62, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_63, negated_conjecture, (r1_tarski(esk3_0,X1)|~r1_tarski(k2_struct_0(esk2_0),X1)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_64, plain, (r1_tarski(k2_struct_0(X1),k2_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[c_0_60, c_0_35])).
cnf(c_0_65, negated_conjecture, (~m1_subset_1(esk3_0,k1_zfmisc_1(k2_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_47]), c_0_37])])).
cnf(c_0_66, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_62, c_0_49])).
cnf(c_0_67, negated_conjecture, (r1_tarski(esk3_0,k2_struct_0(X1))|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_68, negated_conjecture, (~r2_hidden(esk3_0,k1_zfmisc_1(k2_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_69, negated_conjecture, (r2_hidden(esk3_0,k1_zfmisc_1(k2_struct_0(X1)))|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_40, c_0_67])).
cnf(c_0_70, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_36]), c_0_37])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 71
# Proof object clause steps            : 40
# Proof object formula steps           : 31
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 15
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 32
# Processed clauses                    : 135
# ...of these trivial                  : 1
# ...subsumed                          : 31
# ...remaining for further processing  : 103
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 148
# ...of the previous two non-trivial   : 136
# Contextual simplify-reflections      : 6
# Paramodulations                      : 145
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 69
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 50
# Current number of unprocessed clauses: 60
# ...number of literals in the above   : 227
# Current number of archived formulas  : 0
# Current number of archived clauses   : 32
# Clause-clause subsumption calls (NU) : 1575
# Rec. Clause-clause subsumption calls : 521
# Non-unit clause-clause subsumptions  : 29
# Unit Clause-clause subsumption calls : 43
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 11
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4371

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28580 pages
