# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/1778c782f91f45b9b6ed49c847ab43e0.E_input', redefinition_k6_domain_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/1778c782f91f45b9b6ed49c847ab43e0.E_input', t69_enumset1)).
fof(t52_orders_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_orders_2(X1,X2,X3)<=>r2_hidden(X2,k2_orders_2(X1,k6_domain_1(u1_struct_0(X1),X3))))))), file('data/MPTP/proofs/1778c782f91f45b9b6ed49c847ab43e0.E_input', t52_orders_2)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/1778c782f91f45b9b6ed49c847ab43e0.E_input', d1_tarski)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/1778c782f91f45b9b6ed49c847ab43e0.E_input', fc2_struct_0)).
fof(fraenkel_a_2_1_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&(v3_orders_2(X2)&(v4_orders_2(X2)&(v5_orders_2(X2)&l1_orders_2(X2)))))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r2_hidden(X1,a_2_1_orders_2(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r2_hidden(X5,X3)=>r2_orders_2(X2,X4,X5))))))), file('data/MPTP/proofs/1778c782f91f45b9b6ed49c847ab43e0.E_input', fraenkel_a_2_1_orders_2)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/1778c782f91f45b9b6ed49c847ab43e0.E_input', dt_l1_orders_2)).
fof(t35_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v6_orders_2(k6_domain_1(u1_struct_0(X1),X2),X1)&m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/1778c782f91f45b9b6ed49c847ab43e0.E_input', t35_orders_2)).
fof(d13_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k2_orders_2(X1,X2)=a_2_1_orders_2(X1,X2))), file('data/MPTP/proofs/1778c782f91f45b9b6ed49c847ab43e0.E_input', d13_orders_2)).
fof(c_0_9, plain, ![X6, X7]:(v1_xboole_0(X6)|~m1_subset_1(X7,X6)|k6_domain_1(X6,X7)=k1_tarski(X7)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_10, plain, ![X8]:k2_tarski(X8,X8)=k1_tarski(X8), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_11, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_orders_2(X1,X2,X3)<=>r2_hidden(X2,k2_orders_2(X1,k6_domain_1(u1_struct_0(X1),X3)))))))), inference(assume_negation,[status(cth)],[t52_orders_2])).
cnf(c_0_12, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_14, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&((~r2_orders_2(esk1_0,esk2_0,esk3_0)|~r2_hidden(esk2_0,k2_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0))))&(r2_orders_2(esk1_0,esk2_0,esk3_0)|r2_hidden(esk2_0,k2_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])).
fof(c_0_15, plain, ![X12, X13, X14, X15, X16, X17]:(((~r2_hidden(X14,X13)|X14=X12|X13!=k1_tarski(X12))&(X15!=X12|r2_hidden(X15,X13)|X13!=k1_tarski(X12)))&((~r2_hidden(esk4_2(X16,X17),X17)|esk4_2(X16,X17)!=X16|X17=k1_tarski(X16))&(r2_hidden(esk4_2(X16,X17),X17)|esk4_2(X16,X17)=X16|X17=k1_tarski(X16)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_16, plain, ![X19]:(v2_struct_0(X19)|~l1_struct_0(X19)|~v1_xboole_0(u1_struct_0(X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_17, plain, (k2_tarski(X2,X2)=k6_domain_1(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_18, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_19, plain, ![X20, X21, X22, X24, X25]:(((m1_subset_1(esk5_3(X20,X21,X22),u1_struct_0(X21))|~r2_hidden(X20,a_2_1_orders_2(X21,X22))|(v2_struct_0(X21)|(~v3_orders_2(X21)|(~v4_orders_2(X21)|(~v5_orders_2(X21)|~l1_orders_2(X21))))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))))&((X20=esk5_3(X20,X21,X22)|~r2_hidden(X20,a_2_1_orders_2(X21,X22))|(v2_struct_0(X21)|(~v3_orders_2(X21)|(~v4_orders_2(X21)|(~v5_orders_2(X21)|~l1_orders_2(X21))))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))))&(~m1_subset_1(X24,u1_struct_0(X21))|(~r2_hidden(X24,X22)|r2_orders_2(X21,esk5_3(X20,X21,X22),X24))|~r2_hidden(X20,a_2_1_orders_2(X21,X22))|(v2_struct_0(X21)|(~v3_orders_2(X21)|(~v4_orders_2(X21)|(~v5_orders_2(X21)|~l1_orders_2(X21))))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))))))&((m1_subset_1(esk6_4(X20,X21,X22,X25),u1_struct_0(X21))|X20!=X25|~m1_subset_1(X25,u1_struct_0(X21))|r2_hidden(X20,a_2_1_orders_2(X21,X22))|(v2_struct_0(X21)|(~v3_orders_2(X21)|(~v4_orders_2(X21)|(~v5_orders_2(X21)|~l1_orders_2(X21))))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))))&((r2_hidden(esk6_4(X20,X21,X22,X25),X22)|X20!=X25|~m1_subset_1(X25,u1_struct_0(X21))|r2_hidden(X20,a_2_1_orders_2(X21,X22))|(v2_struct_0(X21)|(~v3_orders_2(X21)|(~v4_orders_2(X21)|(~v5_orders_2(X21)|~l1_orders_2(X21))))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))))&(~r2_orders_2(X21,X25,esk6_4(X20,X21,X22,X25))|X20!=X25|~m1_subset_1(X25,u1_struct_0(X21))|r2_hidden(X20,a_2_1_orders_2(X21,X22))|(v2_struct_0(X21)|(~v3_orders_2(X21)|(~v4_orders_2(X21)|(~v5_orders_2(X21)|~l1_orders_2(X21))))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_1_orders_2])])])])])])).
cnf(c_0_20, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, negated_conjecture, (k2_tarski(esk3_0,esk3_0)=k6_domain_1(u1_struct_0(esk1_0),esk3_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_23, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_24, plain, ![X27]:(~l1_orders_2(X27)|l1_struct_0(X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_25, plain, (r2_hidden(esk6_4(X1,X2,X3,X4),X3)|r2_hidden(X1,a_2_1_orders_2(X2,X3))|v2_struct_0(X2)|X1!=X4|~m1_subset_1(X4,u1_struct_0(X2))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_26, plain, ![X28, X29]:((v6_orders_2(k6_domain_1(u1_struct_0(X28),X29),X28)|~m1_subset_1(X29,u1_struct_0(X28))|(v2_struct_0(X28)|(~v3_orders_2(X28)|~l1_orders_2(X28))))&(m1_subset_1(k6_domain_1(u1_struct_0(X28),X29),k1_zfmisc_1(u1_struct_0(X28)))|~m1_subset_1(X29,u1_struct_0(X28))|(v2_struct_0(X28)|(~v3_orders_2(X28)|~l1_orders_2(X28))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t35_orders_2])])])])])).
cnf(c_0_27, plain, (X1=X3|X2!=k2_tarski(X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_20, c_0_13])).
cnf(c_0_28, negated_conjecture, (k2_tarski(esk3_0,esk3_0)=k6_domain_1(u1_struct_0(esk1_0),esk3_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])).
cnf(c_0_29, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_31, plain, (r2_hidden(esk6_4(X1,X2,X3,X1),X3)|r2_hidden(X1,a_2_1_orders_2(X2,X3))|v2_struct_0(X2)|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))), inference(er,[status(thm)],[c_0_25])).
cnf(c_0_32, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_33, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_34, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_35, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_36, plain, (m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_37, plain, ![X30, X31]:(v2_struct_0(X30)|(~v3_orders_2(X30)|(~v4_orders_2(X30)|(~v5_orders_2(X30)|~l1_orders_2(X30))))|(~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))|k2_orders_2(X30,X31)=a_2_1_orders_2(X30,X31))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_orders_2])])])])).
cnf(c_0_38, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_39, plain, (X1=X2|~r2_hidden(X1,k2_tarski(X2,X2))), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_40, negated_conjecture, (k2_tarski(esk3_0,esk3_0)=k6_domain_1(u1_struct_0(esk1_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])])).
cnf(c_0_41, negated_conjecture, (r2_hidden(esk6_4(esk2_0,esk1_0,X1,esk2_0),X1)|r2_hidden(esk2_0,a_2_1_orders_2(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_30]), c_0_33]), c_0_34]), c_0_35])]), c_0_23])).
cnf(c_0_42, negated_conjecture, (m1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_18]), c_0_30]), c_0_35])]), c_0_23])).
cnf(c_0_43, plain, (v2_struct_0(X1)|k2_orders_2(X1,X2)=a_2_1_orders_2(X1,X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_44, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X2)), inference(rw,[status(thm)],[c_0_38, c_0_13])).
cnf(c_0_45, plain, (r2_hidden(X3,a_2_1_orders_2(X1,X4))|v2_struct_0(X1)|~r2_orders_2(X1,X2,esk6_4(X3,X1,X4,X2))|X3!=X2|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_46, negated_conjecture, (X1=esk3_0|~r2_hidden(X1,k6_domain_1(u1_struct_0(esk1_0),esk3_0))), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_47, negated_conjecture, (r2_hidden(esk6_4(esk2_0,esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0),esk2_0),k6_domain_1(u1_struct_0(esk1_0),esk3_0))|r2_hidden(esk2_0,a_2_1_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_48, negated_conjecture, (r2_orders_2(esk1_0,esk2_0,esk3_0)|r2_hidden(esk2_0,k2_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_49, negated_conjecture, (k2_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0))=a_2_1_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_42]), c_0_30]), c_0_33]), c_0_34]), c_0_35])]), c_0_23])).
cnf(c_0_50, plain, (r2_hidden(X1,k2_tarski(X1,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_44])])).
cnf(c_0_51, plain, (r2_hidden(X1,a_2_1_orders_2(X2,X3))|v2_struct_0(X2)|~r2_orders_2(X2,X1,esk6_4(X1,X2,X3,X1))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))), inference(er,[status(thm)],[c_0_45])).
cnf(c_0_52, negated_conjecture, (esk6_4(esk2_0,esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0),esk2_0)=esk3_0|r2_hidden(esk2_0,a_2_1_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_53, negated_conjecture, (r2_hidden(esk2_0,a_2_1_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))|r2_orders_2(esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_54, negated_conjecture, (~r2_orders_2(esk1_0,esk2_0,esk3_0)|~r2_hidden(esk2_0,k2_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_55, plain, (r2_orders_2(X2,esk5_3(X4,X2,X3),X1)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r2_hidden(X4,a_2_1_orders_2(X2,X3))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_56, negated_conjecture, (r2_hidden(esk3_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0))), inference(spm,[status(thm)],[c_0_50, c_0_40])).
cnf(c_0_57, plain, (X1=esk5_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_orders_2(X2,X3))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_58, negated_conjecture, (r2_hidden(esk2_0,a_2_1_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_30]), c_0_33]), c_0_34]), c_0_35]), c_0_42]), c_0_32])]), c_0_23]), c_0_53])).
cnf(c_0_59, negated_conjecture, (~r2_hidden(esk2_0,a_2_1_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))|~r2_orders_2(esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_54, c_0_49])).
cnf(c_0_60, negated_conjecture, (r2_orders_2(X1,esk5_3(X2,X1,k6_domain_1(u1_struct_0(esk1_0),esk3_0)),esk3_0)|v2_struct_0(X1)|~r2_hidden(X2,a_2_1_orders_2(X1,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk3_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk3_0,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_61, negated_conjecture, (esk5_3(esk2_0,esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0))=esk2_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_30]), c_0_33]), c_0_34]), c_0_35]), c_0_42])]), c_0_23])).
cnf(c_0_62, negated_conjecture, (~r2_orders_2(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_58])])).
cnf(c_0_63, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_58]), c_0_61]), c_0_30]), c_0_33]), c_0_34]), c_0_35]), c_0_42]), c_0_18])]), c_0_62]), c_0_23]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 64
# Proof object clause steps            : 45
# Proof object formula steps           : 19
# Proof object conjectures             : 29
# Proof object clause conjectures      : 26
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 57
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 25
# Processed clauses                    : 118
# ...of these trivial                  : 1
# ...subsumed                          : 6
# ...remaining for further processing  : 111
# Other redundant clauses eliminated   : 15
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 11
# Generated clauses                    : 149
# ...of the previous two non-trivial   : 108
# Contextual simplify-reflections      : 2
# Paramodulations                      : 129
# Factorizations                       : 6
# Equation resolutions                 : 15
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 69
#    Positive orientable unit clauses  : 19
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 48
# Current number of unprocessed clauses: 38
# ...number of literals in the above   : 187
# Current number of archived formulas  : 0
# Current number of archived clauses   : 38
# Clause-clause subsumption calls (NU) : 676
# Rec. Clause-clause subsumption calls : 113
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6671

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28640 pages
