# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d16_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r3_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X2,X4)))))), file('data/MPTP/proofs/17ef801e17034486af786b150222043c.E_input', d16_lattice3)).
fof(fraenkel_a_2_1_lattice3, axiom, ![X1, X2, X3]:((~(v2_struct_0(X2))&l3_lattices(X2))=>(r2_hidden(X1,a_2_1_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&r3_lattice3(X2,X4,X3))))), file('data/MPTP/proofs/17ef801e17034486af786b150222043c.E_input', fraenkel_a_2_1_lattice3)).
fof(d17_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r4_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X4,X2)))))), file('data/MPTP/proofs/17ef801e17034486af786b150222043c.E_input', d17_lattice3)).
fof(t41_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((r2_hidden(X2,X3)&r4_lattice3(X1,X2,X3))=>k15_lattice3(X1,X3)=X2))), file('data/MPTP/proofs/17ef801e17034486af786b150222043c.E_input', t41_lattice3)).
fof(t42_lattice3, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((r2_hidden(X2,X3)&r3_lattice3(X1,X2,X3))=>k16_lattice3(X1,X3)=X2))), file('data/MPTP/proofs/17ef801e17034486af786b150222043c.E_input', t42_lattice3)).
fof(t34_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(X2=k16_lattice3(X1,X3)<=>(r3_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_lattice3(X1,X4,X3)=>r3_lattices(X1,X4,X2))))))), file('data/MPTP/proofs/17ef801e17034486af786b150222043c.E_input', t34_lattice3)).
fof(t38_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r2_hidden(X2,X3)=>(r3_lattices(X1,X2,k15_lattice3(X1,X3))&r3_lattices(X1,k16_lattice3(X1,X3),X2))))), file('data/MPTP/proofs/17ef801e17034486af786b150222043c.E_input', t38_lattice3)).
fof(c_0_7, plain, ![X5, X6, X7, X8, X9]:((~r3_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X7)|r1_lattices(X5,X6,X8)))|~m1_subset_1(X6,u1_struct_0(X5))|(v2_struct_0(X5)|~l3_lattices(X5)))&((m1_subset_1(esk1_3(X5,X6,X9),u1_struct_0(X5))|r3_lattice3(X5,X6,X9)|~m1_subset_1(X6,u1_struct_0(X5))|(v2_struct_0(X5)|~l3_lattices(X5)))&((r2_hidden(esk1_3(X5,X6,X9),X9)|r3_lattice3(X5,X6,X9)|~m1_subset_1(X6,u1_struct_0(X5))|(v2_struct_0(X5)|~l3_lattices(X5)))&(~r1_lattices(X5,X6,esk1_3(X5,X6,X9))|r3_lattice3(X5,X6,X9)|~m1_subset_1(X6,u1_struct_0(X5))|(v2_struct_0(X5)|~l3_lattices(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattice3])])])])])])])).
fof(c_0_8, plain, ![X11, X12, X13, X15, X16]:(((m1_subset_1(esk2_3(X11,X12,X13),u1_struct_0(X12))|~r2_hidden(X11,a_2_1_lattice3(X12,X13))|(v2_struct_0(X12)|~l3_lattices(X12)))&((X11=esk2_3(X11,X12,X13)|~r2_hidden(X11,a_2_1_lattice3(X12,X13))|(v2_struct_0(X12)|~l3_lattices(X12)))&(r3_lattice3(X12,esk2_3(X11,X12,X13),X13)|~r2_hidden(X11,a_2_1_lattice3(X12,X13))|(v2_struct_0(X12)|~l3_lattices(X12)))))&(~m1_subset_1(X16,u1_struct_0(X12))|(X11!=X16|~r3_lattice3(X12,X16,X15))|r2_hidden(X11,a_2_1_lattice3(X12,X15))|(v2_struct_0(X12)|~l3_lattices(X12)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_1_lattice3])])])])])])])).
cnf(c_0_9, plain, (r1_lattices(X1,X2,X4)|v2_struct_0(X1)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10, plain, (r3_lattice3(X1,esk2_3(X2,X1,X3),X3)|v2_struct_0(X1)|~r2_hidden(X2,a_2_1_lattice3(X1,X3))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11, plain, (m1_subset_1(esk2_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_lattice3(X2,X3))|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (r1_lattices(X1,esk2_3(X2,X1,X3),X4)|v2_struct_0(X1)|~r2_hidden(X2,a_2_1_lattice3(X1,X3))|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11])).
cnf(c_0_13, plain, (X1=esk2_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_lattice3(X2,X3))|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_14, plain, ![X17, X18, X19, X20, X21]:((~r4_lattice3(X17,X18,X19)|(~m1_subset_1(X20,u1_struct_0(X17))|(~r2_hidden(X20,X19)|r1_lattices(X17,X20,X18)))|~m1_subset_1(X18,u1_struct_0(X17))|(v2_struct_0(X17)|~l3_lattices(X17)))&((m1_subset_1(esk3_3(X17,X18,X21),u1_struct_0(X17))|r4_lattice3(X17,X18,X21)|~m1_subset_1(X18,u1_struct_0(X17))|(v2_struct_0(X17)|~l3_lattices(X17)))&((r2_hidden(esk3_3(X17,X18,X21),X21)|r4_lattice3(X17,X18,X21)|~m1_subset_1(X18,u1_struct_0(X17))|(v2_struct_0(X17)|~l3_lattices(X17)))&(~r1_lattices(X17,esk3_3(X17,X18,X21),X18)|r4_lattice3(X17,X18,X21)|~m1_subset_1(X18,u1_struct_0(X17))|(v2_struct_0(X17)|~l3_lattices(X17)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_lattice3])])])])])])])).
cnf(c_0_15, plain, (r1_lattices(X1,X2,X3)|v2_struct_0(X1)|~r2_hidden(X2,a_2_1_lattice3(X1,X4))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_16, plain, (r2_hidden(esk3_3(X1,X2,X3),X3)|r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_17, plain, ![X23, X24, X25]:(v2_struct_0(X23)|(~v10_lattices(X23)|(~v4_lattice3(X23)|~l3_lattices(X23)))|(~m1_subset_1(X24,u1_struct_0(X23))|(~r2_hidden(X24,X25)|~r4_lattice3(X23,X24,X25)|k15_lattice3(X23,X25)=X24))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t41_lattice3])])])])).
cnf(c_0_18, plain, (r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r1_lattices(X1,esk3_3(X1,X2,X3),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (r4_lattice3(X1,X2,a_2_1_lattice3(X3,X4))|r1_lattices(X3,esk3_3(X1,X2,a_2_1_lattice3(X3,X4)),X5)|v2_struct_0(X1)|v2_struct_0(X3)|~r2_hidden(X5,X4)|~m1_subset_1(X5,u1_struct_0(X3))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X3)|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_20, plain, (v2_struct_0(X1)|k15_lattice3(X1,X3)=X2|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X2,X3)|~r4_lattice3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21, plain, (r4_lattice3(X1,X2,a_2_1_lattice3(X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_22, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_lattice3(X2,X3))|~l3_lattices(X2)), inference(spm,[status(thm)],[c_0_11, c_0_13])).
cnf(c_0_23, plain, (r2_hidden(X3,a_2_1_lattice3(X2,X4))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~r3_lattice3(X2,X1,X4)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_24, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((r2_hidden(X2,X3)&r3_lattice3(X1,X2,X3))=>k16_lattice3(X1,X3)=X2)))), inference(assume_negation,[status(cth)],[t42_lattice3])).
fof(c_0_25, plain, ![X29, X30, X31, X32, X33]:(((r3_lattice3(X29,X30,X31)|X30!=k16_lattice3(X29,X31)|~m1_subset_1(X30,u1_struct_0(X29))|(v2_struct_0(X29)|(~v10_lattices(X29)|(~v4_lattice3(X29)|~l3_lattices(X29)))))&(~m1_subset_1(X32,u1_struct_0(X29))|(~r3_lattice3(X29,X32,X31)|r3_lattices(X29,X32,X30))|X30!=k16_lattice3(X29,X31)|~m1_subset_1(X30,u1_struct_0(X29))|(v2_struct_0(X29)|(~v10_lattices(X29)|(~v4_lattice3(X29)|~l3_lattices(X29))))))&((m1_subset_1(esk7_3(X29,X30,X33),u1_struct_0(X29))|~r3_lattice3(X29,X30,X33)|X30=k16_lattice3(X29,X33)|~m1_subset_1(X30,u1_struct_0(X29))|(v2_struct_0(X29)|(~v10_lattices(X29)|(~v4_lattice3(X29)|~l3_lattices(X29)))))&((r3_lattice3(X29,esk7_3(X29,X30,X33),X33)|~r3_lattice3(X29,X30,X33)|X30=k16_lattice3(X29,X33)|~m1_subset_1(X30,u1_struct_0(X29))|(v2_struct_0(X29)|(~v10_lattices(X29)|(~v4_lattice3(X29)|~l3_lattices(X29)))))&(~r3_lattices(X29,esk7_3(X29,X30,X33),X30)|~r3_lattice3(X29,X30,X33)|X30=k16_lattice3(X29,X33)|~m1_subset_1(X30,u1_struct_0(X29))|(v2_struct_0(X29)|(~v10_lattices(X29)|(~v4_lattice3(X29)|~l3_lattices(X29)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_lattice3])])])])])])])).
fof(c_0_26, plain, ![X35, X36, X37]:((r3_lattices(X35,X36,k15_lattice3(X35,X37))|~r2_hidden(X36,X37)|~m1_subset_1(X36,u1_struct_0(X35))|(v2_struct_0(X35)|(~v10_lattices(X35)|(~v4_lattice3(X35)|~l3_lattices(X35)))))&(r3_lattices(X35,k16_lattice3(X35,X37),X36)|~r2_hidden(X36,X37)|~m1_subset_1(X36,u1_struct_0(X35))|(v2_struct_0(X35)|(~v10_lattices(X35)|(~v4_lattice3(X35)|~l3_lattices(X35)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t38_lattice3])])])])])).
cnf(c_0_27, plain, (k15_lattice3(X1,a_2_1_lattice3(X1,X2))=X3|v2_struct_0(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(X3,a_2_1_lattice3(X1,X2))|~r2_hidden(X3,X2)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])).
cnf(c_0_28, plain, (r2_hidden(X1,a_2_1_lattice3(X2,X3))|v2_struct_0(X2)|~r3_lattice3(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)), inference(er,[status(thm)],[c_0_23])).
fof(c_0_29, negated_conjecture, ((~v2_struct_0(esk4_0)&(v10_lattices(esk4_0)&(v4_lattice3(esk4_0)&l3_lattices(esk4_0))))&(m1_subset_1(esk5_0,u1_struct_0(esk4_0))&((r2_hidden(esk5_0,esk6_0)&r3_lattice3(esk4_0,esk5_0,esk6_0))&k16_lattice3(esk4_0,esk6_0)!=esk5_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_24])])])])).
cnf(c_0_30, plain, (X2=k16_lattice3(X1,X3)|v2_struct_0(X1)|~r3_lattices(X1,esk7_3(X1,X2,X3),X2)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, plain, (r3_lattices(X1,X2,k15_lattice3(X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|X2=k16_lattice3(X1,X3)|v2_struct_0(X1)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (k15_lattice3(X1,a_2_1_lattice3(X1,X2))=X3|v2_struct_0(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(X3,X2)|~r3_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34, negated_conjecture, (r3_lattice3(esk4_0,esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_35, negated_conjecture, (v4_lattice3(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, negated_conjecture, (v10_lattices(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, negated_conjecture, (r2_hidden(esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, negated_conjecture, (l3_lattices(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_40, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_41, plain, (k15_lattice3(X1,X2)=k16_lattice3(X1,X3)|v2_struct_0(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(esk7_3(X1,k15_lattice3(X1,X2),X3),X2)|~r3_lattice3(X1,k15_lattice3(X1,X2),X3)|~m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])).
cnf(c_0_42, negated_conjecture, (k15_lattice3(esk4_0,a_2_1_lattice3(esk4_0,esk6_0))=esk5_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39])]), c_0_40])).
cnf(c_0_43, negated_conjecture, (esk5_0=k16_lattice3(esk4_0,X1)|~r2_hidden(esk7_3(esk4_0,esk5_0,X1),a_2_1_lattice3(esk4_0,esk6_0))|~r3_lattice3(esk4_0,esk5_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_35]), c_0_36]), c_0_38]), c_0_39])]), c_0_40])).
cnf(c_0_44, negated_conjecture, (esk5_0=k16_lattice3(esk4_0,X1)|~r3_lattice3(esk4_0,esk7_3(esk4_0,esk5_0,X1),esk6_0)|~r3_lattice3(esk4_0,esk5_0,X1)|~m1_subset_1(esk7_3(esk4_0,esk5_0,X1),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_28]), c_0_39])]), c_0_40])).
cnf(c_0_45, plain, (r3_lattice3(X1,esk7_3(X1,X2,X3),X3)|X2=k16_lattice3(X1,X3)|v2_struct_0(X1)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_46, negated_conjecture, (k16_lattice3(esk4_0,esk6_0)!=esk5_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_47, negated_conjecture, (~m1_subset_1(esk7_3(esk4_0,esk5_0,esk6_0),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_34]), c_0_35]), c_0_36]), c_0_38]), c_0_39])]), c_0_46]), c_0_40])).
cnf(c_0_48, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_32]), c_0_35]), c_0_36]), c_0_34]), c_0_38]), c_0_39])]), c_0_46]), c_0_40]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 49
# Proof object clause steps            : 34
# Proof object formula steps           : 15
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 36
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 28
# Processed clauses                    : 97
# ...of these trivial                  : 0
# ...subsumed                          : 4
# ...remaining for further processing  : 93
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 71
# ...of the previous two non-trivial   : 66
# Contextual simplify-reflections      : 7
# Paramodulations                      : 68
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 64
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 52
# Current number of unprocessed clauses: 25
# ...number of literals in the above   : 224
# Current number of archived formulas  : 0
# Current number of archived clauses   : 28
# Clause-clause subsumption calls (NU) : 986
# Rec. Clause-clause subsumption calls : 53
# Non-unit clause-clause subsumptions  : 11
# Unit Clause-clause subsumption calls : 20
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5463

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.000 s
# Total time               : 0.013 s
# Maximum resident set size: 28748 pages
