# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t74_tops_2, conjecture, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v3_tops_2(X3,X1,X2)<=>(k1_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X1)&(k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X2)&(v2_funct_1(X3)&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,k2_pre_topc(X1,X4))=k2_pre_topc(X2,k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4)))))))))), file('data/MPTP/proofs/1819d8422cbf42ee8b176653c59ceba8.E_input', t74_tops_2)).
fof(t57_tops_2, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v5_pre_topc(X3,X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>r1_tarski(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,k2_pre_topc(X1,X4)),k2_pre_topc(X2,k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4)))))))), file('data/MPTP/proofs/1819d8422cbf42ee8b176653c59ceba8.E_input', t57_tops_2)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/1819d8422cbf42ee8b176653c59ceba8.E_input', d10_xboole_0)).
fof(dt_k2_tops_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(v1_funct_1(k2_tops_2(X1,X2,X3))&(v1_funct_2(k2_tops_2(X1,X2,X3),X2,X1)&m1_subset_1(k2_tops_2(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1)))))), file('data/MPTP/proofs/1819d8422cbf42ee8b176653c59ceba8.E_input', dt_k2_tops_2)).
fof(t67_tops_2, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(l1_struct_0(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>((k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X2)&v2_funct_1(X3))=>k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4)=k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X4)))))), file('data/MPTP/proofs/1819d8422cbf42ee8b176653c59ceba8.E_input', t67_tops_2)).
fof(d12_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v5_pre_topc(X3,X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(v4_pre_topc(X4,X2)=>v4_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X1))))))), file('data/MPTP/proofs/1819d8422cbf42ee8b176653c59ceba8.E_input', d12_pre_topc)).
fof(d5_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v3_tops_2(X3,X1,X2)<=>(k1_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X1)&(k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X2)&(v2_funct_1(X3)&(v5_pre_topc(X3,X1,X2)&v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1))))))))), file('data/MPTP/proofs/1819d8422cbf42ee8b176653c59ceba8.E_input', d5_tops_2)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/1819d8422cbf42ee8b176653c59ceba8.E_input', dt_l1_pre_topc)).
fof(t52_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v4_pre_topc(X2,X1)=>k2_pre_topc(X1,X2)=X2)&((v2_pre_topc(X1)&k2_pre_topc(X1,X2)=X2)=>v4_pre_topc(X2,X1))))), file('data/MPTP/proofs/1819d8422cbf42ee8b176653c59ceba8.E_input', t52_pre_topc)).
fof(dt_k7_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k7_relset_1(X1,X2,X3,X4),k1_zfmisc_1(X2))), file('data/MPTP/proofs/1819d8422cbf42ee8b176653c59ceba8.E_input', dt_k7_relset_1)).
fof(t70_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:((~(v2_struct_0(X2))&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v3_tops_2(X3,X1,X2)=>v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1))))), file('data/MPTP/proofs/1819d8422cbf42ee8b176653c59ceba8.E_input', t70_tops_2)).
fof(dt_k2_pre_topc, axiom, ![X1, X2]:((l1_pre_topc(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k2_pre_topc(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/1819d8422cbf42ee8b176653c59ceba8.E_input', dt_k2_pre_topc)).
fof(t73_tops_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((v2_pre_topc(X2)&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v3_tops_2(X3,X1,X2)<=>(k1_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X1)&(k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X2)&(v2_funct_1(X3)&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,k2_pre_topc(X2,X4))=k2_pre_topc(X1,k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4)))))))))), file('data/MPTP/proofs/1819d8422cbf42ee8b176653c59ceba8.E_input', t73_tops_2)).
fof(c_0_13, negated_conjecture, ~(![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v3_tops_2(X3,X1,X2)<=>(k1_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X1)&(k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X2)&(v2_funct_1(X3)&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,k2_pre_topc(X1,X4))=k2_pre_topc(X2,k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4))))))))))), inference(assume_negation,[status(cth)],[t74_tops_2])).
fof(c_0_14, plain, ![X39, X40, X41, X42]:((~v5_pre_topc(X41,X39,X40)|(~m1_subset_1(X42,k1_zfmisc_1(u1_struct_0(X39)))|r1_tarski(k7_relset_1(u1_struct_0(X39),u1_struct_0(X40),X41,k2_pre_topc(X39,X42)),k2_pre_topc(X40,k7_relset_1(u1_struct_0(X39),u1_struct_0(X40),X41,X42))))|(~v1_funct_1(X41)|(~v1_funct_2(X41,u1_struct_0(X39),u1_struct_0(X40))|~m1_subset_1(X41,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X39),u1_struct_0(X40))))))|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40)))|(~v2_pre_topc(X39)|~l1_pre_topc(X39)))&((m1_subset_1(esk6_3(X39,X40,X41),k1_zfmisc_1(u1_struct_0(X39)))|v5_pre_topc(X41,X39,X40)|(~v1_funct_1(X41)|(~v1_funct_2(X41,u1_struct_0(X39),u1_struct_0(X40))|~m1_subset_1(X41,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X39),u1_struct_0(X40))))))|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40)))|(~v2_pre_topc(X39)|~l1_pre_topc(X39)))&(~r1_tarski(k7_relset_1(u1_struct_0(X39),u1_struct_0(X40),X41,k2_pre_topc(X39,esk6_3(X39,X40,X41))),k2_pre_topc(X40,k7_relset_1(u1_struct_0(X39),u1_struct_0(X40),X41,esk6_3(X39,X40,X41))))|v5_pre_topc(X41,X39,X40)|(~v1_funct_1(X41)|(~v1_funct_2(X41,u1_struct_0(X39),u1_struct_0(X40))|~m1_subset_1(X41,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X39),u1_struct_0(X40))))))|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40)))|(~v2_pre_topc(X39)|~l1_pre_topc(X39))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t57_tops_2])])])])])])).
fof(c_0_15, negated_conjecture, ![X9]:((v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0))&((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))))&(((m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_funct_1(esk3_0)|k2_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|k1_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))&(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,k2_pre_topc(esk1_0,esk4_0))!=k2_pre_topc(esk2_0,k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_0))|~v2_funct_1(esk3_0)|k2_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|k1_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)))&((k1_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0))&((k2_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0))&((v2_funct_1(esk3_0)|v3_tops_2(esk3_0,esk1_0,esk2_0))&(~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(esk1_0)))|k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,k2_pre_topc(esk1_0,X9))=k2_pre_topc(esk2_0,k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X9))|v3_tops_2(esk3_0,esk1_0,esk2_0))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_13])])])])])])).
cnf(c_0_16, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|v5_pre_topc(X3,X1,X2)|v2_struct_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_17, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_18, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, negated_conjecture, (v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_25, plain, ![X31, X32]:(((r1_tarski(X31,X32)|X31!=X32)&(r1_tarski(X32,X31)|X31!=X32))&(~r1_tarski(X31,X32)|~r1_tarski(X32,X31)|X31=X32)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_26, plain, ![X18, X19, X20]:((v1_funct_1(k2_tops_2(X18,X19,X20))|(~v1_funct_1(X20)|(~v1_funct_2(X20,X18,X19)|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(X18,X19))))))&((v1_funct_2(k2_tops_2(X18,X19,X20),X19,X18)|(~v1_funct_1(X20)|(~v1_funct_2(X20,X18,X19)|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(X18,X19))))))&(m1_subset_1(k2_tops_2(X18,X19,X20),k1_zfmisc_1(k2_zfmisc_1(X19,X18)))|(~v1_funct_1(X20)|(~v1_funct_2(X20,X18,X19)|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(X18,X19)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_tops_2])])])).
cnf(c_0_27, negated_conjecture, (k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,k2_pre_topc(esk1_0,X1))=k2_pre_topc(esk2_0,k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1))|v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_28, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,esk2_0)|m1_subset_1(esk6_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23])]), c_0_24])).
cnf(c_0_29, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_30, plain, ![X14, X15, X16, X17]:(~l1_struct_0(X14)|(~l1_struct_0(X15)|(~v1_funct_1(X16)|(~v1_funct_2(X16,u1_struct_0(X14),u1_struct_0(X15))|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X14),u1_struct_0(X15)))))|(~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X14)))|(k2_relset_1(u1_struct_0(X14),u1_struct_0(X15),X16)!=k2_struct_0(X15)|~v2_funct_1(X16)|k7_relset_1(u1_struct_0(X14),u1_struct_0(X15),X16,X17)=k8_relset_1(u1_struct_0(X15),u1_struct_0(X14),k2_tops_2(u1_struct_0(X14),u1_struct_0(X15),X16),X17)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t67_tops_2])])])).
fof(c_0_31, plain, ![X21, X22, X23, X24]:((~v5_pre_topc(X23,X21,X22)|(~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X22)))|(~v4_pre_topc(X24,X22)|v4_pre_topc(k8_relset_1(u1_struct_0(X21),u1_struct_0(X22),X23,X24),X21)))|(~v1_funct_1(X23)|(~v1_funct_2(X23,u1_struct_0(X21),u1_struct_0(X22))|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X21),u1_struct_0(X22))))))|~l1_pre_topc(X22)|~l1_pre_topc(X21))&((m1_subset_1(esk5_3(X21,X22,X23),k1_zfmisc_1(u1_struct_0(X22)))|v5_pre_topc(X23,X21,X22)|(~v1_funct_1(X23)|(~v1_funct_2(X23,u1_struct_0(X21),u1_struct_0(X22))|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X21),u1_struct_0(X22))))))|~l1_pre_topc(X22)|~l1_pre_topc(X21))&((v4_pre_topc(esk5_3(X21,X22,X23),X22)|v5_pre_topc(X23,X21,X22)|(~v1_funct_1(X23)|(~v1_funct_2(X23,u1_struct_0(X21),u1_struct_0(X22))|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X21),u1_struct_0(X22))))))|~l1_pre_topc(X22)|~l1_pre_topc(X21))&(~v4_pre_topc(k8_relset_1(u1_struct_0(X21),u1_struct_0(X22),X23,esk5_3(X21,X22,X23)),X21)|v5_pre_topc(X23,X21,X22)|(~v1_funct_1(X23)|(~v1_funct_2(X23,u1_struct_0(X21),u1_struct_0(X22))|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X21),u1_struct_0(X22))))))|~l1_pre_topc(X22)|~l1_pre_topc(X21))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_pre_topc])])])])])).
cnf(c_0_32, plain, (v1_funct_2(k2_tops_2(X1,X2,X3),X2,X1)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, plain, (m1_subset_1(k2_tops_2(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (v1_funct_1(k2_tops_2(X1,X2,X3))|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_35, plain, ![X26, X27, X28]:(((k1_relset_1(u1_struct_0(X26),u1_struct_0(X27),X28)=k2_struct_0(X26)|~v3_tops_2(X28,X26,X27)|(~v1_funct_1(X28)|(~v1_funct_2(X28,u1_struct_0(X26),u1_struct_0(X27))|~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X26),u1_struct_0(X27))))))|~l1_pre_topc(X27)|~l1_pre_topc(X26))&((k2_relset_1(u1_struct_0(X26),u1_struct_0(X27),X28)=k2_struct_0(X27)|~v3_tops_2(X28,X26,X27)|(~v1_funct_1(X28)|(~v1_funct_2(X28,u1_struct_0(X26),u1_struct_0(X27))|~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X26),u1_struct_0(X27))))))|~l1_pre_topc(X27)|~l1_pre_topc(X26))&((v2_funct_1(X28)|~v3_tops_2(X28,X26,X27)|(~v1_funct_1(X28)|(~v1_funct_2(X28,u1_struct_0(X26),u1_struct_0(X27))|~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X26),u1_struct_0(X27))))))|~l1_pre_topc(X27)|~l1_pre_topc(X26))&((v5_pre_topc(X28,X26,X27)|~v3_tops_2(X28,X26,X27)|(~v1_funct_1(X28)|(~v1_funct_2(X28,u1_struct_0(X26),u1_struct_0(X27))|~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X26),u1_struct_0(X27))))))|~l1_pre_topc(X27)|~l1_pre_topc(X26))&(v5_pre_topc(k2_tops_2(u1_struct_0(X26),u1_struct_0(X27),X28),X27,X26)|~v3_tops_2(X28,X26,X27)|(~v1_funct_1(X28)|(~v1_funct_2(X28,u1_struct_0(X26),u1_struct_0(X27))|~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X26),u1_struct_0(X27))))))|~l1_pre_topc(X27)|~l1_pre_topc(X26))))))&(k1_relset_1(u1_struct_0(X26),u1_struct_0(X27),X28)!=k2_struct_0(X26)|(k2_relset_1(u1_struct_0(X26),u1_struct_0(X27),X28)!=k2_struct_0(X27)|(~v2_funct_1(X28)|(~v5_pre_topc(X28,X26,X27)|~v5_pre_topc(k2_tops_2(u1_struct_0(X26),u1_struct_0(X27),X28),X27,X26))))|v3_tops_2(X28,X26,X27)|(~v1_funct_1(X28)|(~v1_funct_2(X28,u1_struct_0(X26),u1_struct_0(X27))|~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X26),u1_struct_0(X27))))))|~l1_pre_topc(X27)|~l1_pre_topc(X26))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_tops_2])])])])).
cnf(c_0_36, plain, (v5_pre_topc(X3,X1,X2)|v2_struct_0(X2)|~r1_tarski(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,k2_pre_topc(X1,esk6_3(X1,X2,X3))),k2_pre_topc(X2,k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk6_3(X1,X2,X3))))|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_37, negated_conjecture, (k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,k2_pre_topc(esk1_0,esk6_3(esk1_0,esk2_0,esk3_0)))=k2_pre_topc(esk2_0,k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk6_3(esk1_0,esk2_0,esk3_0)))|v5_pre_topc(esk3_0,esk1_0,esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_38, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_29])).
cnf(c_0_39, plain, (k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4)=k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X4)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v2_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40, negated_conjecture, (k2_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_41, negated_conjecture, (v2_funct_1(esk3_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_42, plain, ![X38]:(~l1_pre_topc(X38)|l1_struct_0(X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_43, plain, (m1_subset_1(esk5_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|v5_pre_topc(X3,X1,X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_44, negated_conjecture, (v1_funct_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_17]), c_0_18]), c_0_19])])).
cnf(c_0_45, negated_conjecture, (m1_subset_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_17]), c_0_18]), c_0_19])])).
cnf(c_0_46, negated_conjecture, (v1_funct_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_17]), c_0_18]), c_0_19])])).
cnf(c_0_47, plain, (v3_tops_2(X3,X1,X2)|k1_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v2_funct_1(X3)|~v5_pre_topc(X3,X1,X2)|~v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_48, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_49, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_18]), c_0_17]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23])]), c_0_24])).
fof(c_0_50, plain, ![X29, X30]:((~v4_pre_topc(X30,X29)|k2_pre_topc(X29,X30)=X30|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))|~l1_pre_topc(X29))&(~v2_pre_topc(X29)|k2_pre_topc(X29,X30)!=X30|v4_pre_topc(X30,X29)|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))|~l1_pre_topc(X29))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t52_pre_topc])])])])).
cnf(c_0_51, plain, (v4_pre_topc(esk5_3(X1,X2,X3),X2)|v5_pre_topc(X3,X1,X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_52, negated_conjecture, (k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),X1)=k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~l1_struct_0(esk2_0)|~l1_struct_0(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_18]), c_0_17]), c_0_19])]), c_0_41])).
cnf(c_0_53, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_54, negated_conjecture, (v5_pre_topc(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk2_0,esk1_0)|m1_subset_1(esk5_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_46]), c_0_21]), c_0_20])])).
cnf(c_0_55, negated_conjecture, (v3_tops_2(esk3_0,esk1_0,esk2_0)|~v5_pre_topc(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_18]), c_0_17]), c_0_19]), c_0_20]), c_0_21])]), c_0_41]), c_0_49]), c_0_40])).
cnf(c_0_56, plain, (k2_pre_topc(X2,X1)=X1|~v4_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_57, negated_conjecture, (v4_pre_topc(esk5_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)),esk1_0)|v5_pre_topc(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_44]), c_0_45]), c_0_46]), c_0_21]), c_0_20])])).
fof(c_0_58, plain, ![X10, X11, X12, X13]:(~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(X10,X11)))|m1_subset_1(k7_relset_1(X10,X11,X12,X13),k1_zfmisc_1(X11))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_relset_1])])).
cnf(c_0_59, negated_conjecture, (k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),X1)=k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~l1_struct_0(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_20])])).
cnf(c_0_60, negated_conjecture, (k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,k2_pre_topc(esk1_0,esk5_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0))))=k2_pre_topc(esk2_0,k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk5_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0))))|v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_54]), c_0_55])).
cnf(c_0_61, negated_conjecture, (k2_pre_topc(esk1_0,esk5_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)))=esk5_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0))|v5_pre_topc(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_21])]), c_0_54])).
cnf(c_0_62, plain, (m1_subset_1(k7_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_63, negated_conjecture, (k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),X1)=k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_53]), c_0_21])])).
cnf(c_0_64, plain, (v4_pre_topc(X2,X1)|~v2_pre_topc(X1)|k2_pre_topc(X1,X2)!=X2|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_65, negated_conjecture, (k2_pre_topc(esk2_0,k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk5_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0))))=k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk5_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)))|v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_55])).
cnf(c_0_66, negated_conjecture, (m1_subset_1(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_62, c_0_18])).
cnf(c_0_67, plain, (v5_pre_topc(X3,X1,X2)|~v4_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk5_3(X1,X2,X3)),X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_68, negated_conjecture, (k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk5_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)))=k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk5_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)))|v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_54]), c_0_55])).
cnf(c_0_69, negated_conjecture, (v4_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk5_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0))),esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_66]), c_0_20]), c_0_22])])).
cnf(c_0_70, plain, (k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X2)|~v3_tops_2(X3,X1,X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_71, negated_conjecture, (v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_45]), c_0_44]), c_0_46]), c_0_21]), c_0_20])]), c_0_69]), c_0_55])).
cnf(c_0_72, plain, (v2_funct_1(X1)|~v3_tops_2(X1,X2,X3)|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(X3)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_73, negated_conjecture, (k2_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_18]), c_0_17]), c_0_19]), c_0_20]), c_0_21])])).
cnf(c_0_74, negated_conjecture, (v2_funct_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_71]), c_0_18]), c_0_17]), c_0_19]), c_0_20]), c_0_21])])).
cnf(c_0_75, plain, (k1_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X1)|~v3_tops_2(X3,X1,X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_76, plain, ![X35, X36, X37]:(~l1_pre_topc(X35)|(v2_struct_0(X36)|~l1_pre_topc(X36)|(~v1_funct_1(X37)|(~v1_funct_2(X37,u1_struct_0(X35),u1_struct_0(X36))|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X35),u1_struct_0(X36)))))|(~v3_tops_2(X37,X35,X36)|v3_tops_2(k2_tops_2(u1_struct_0(X35),u1_struct_0(X36),X37),X36,X35))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t70_tops_2])])])])).
cnf(c_0_77, negated_conjecture, (k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),X1)=k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1)|~l1_struct_0(esk2_0)|~l1_struct_0(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_73]), c_0_74]), c_0_18]), c_0_17]), c_0_19])])).
fof(c_0_78, plain, ![X33, X34]:(~l1_pre_topc(X33)|~m1_subset_1(X34,k1_zfmisc_1(u1_struct_0(X33)))|m1_subset_1(k2_pre_topc(X33,X34),k1_zfmisc_1(u1_struct_0(X33)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_pre_topc])])).
cnf(c_0_79, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_funct_1(esk3_0)|k2_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|k1_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_80, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_71]), c_0_18]), c_0_17]), c_0_19]), c_0_20]), c_0_21])])).
fof(c_0_81, plain, ![X44, X45, X46, X47]:(((k1_relset_1(u1_struct_0(X44),u1_struct_0(X45),X46)=k2_struct_0(X44)|~v3_tops_2(X46,X44,X45)|(~v1_funct_1(X46)|(~v1_funct_2(X46,u1_struct_0(X44),u1_struct_0(X45))|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X44),u1_struct_0(X45))))))|(~v2_pre_topc(X45)|~l1_pre_topc(X45))|(v2_struct_0(X44)|(~v2_pre_topc(X44)|~l1_pre_topc(X44))))&((k2_relset_1(u1_struct_0(X44),u1_struct_0(X45),X46)=k2_struct_0(X45)|~v3_tops_2(X46,X44,X45)|(~v1_funct_1(X46)|(~v1_funct_2(X46,u1_struct_0(X44),u1_struct_0(X45))|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X44),u1_struct_0(X45))))))|(~v2_pre_topc(X45)|~l1_pre_topc(X45))|(v2_struct_0(X44)|(~v2_pre_topc(X44)|~l1_pre_topc(X44))))&((v2_funct_1(X46)|~v3_tops_2(X46,X44,X45)|(~v1_funct_1(X46)|(~v1_funct_2(X46,u1_struct_0(X44),u1_struct_0(X45))|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X44),u1_struct_0(X45))))))|(~v2_pre_topc(X45)|~l1_pre_topc(X45))|(v2_struct_0(X44)|(~v2_pre_topc(X44)|~l1_pre_topc(X44))))&(~m1_subset_1(X47,k1_zfmisc_1(u1_struct_0(X45)))|k8_relset_1(u1_struct_0(X44),u1_struct_0(X45),X46,k2_pre_topc(X45,X47))=k2_pre_topc(X44,k8_relset_1(u1_struct_0(X44),u1_struct_0(X45),X46,X47))|~v3_tops_2(X46,X44,X45)|(~v1_funct_1(X46)|(~v1_funct_2(X46,u1_struct_0(X44),u1_struct_0(X45))|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X44),u1_struct_0(X45))))))|(~v2_pre_topc(X45)|~l1_pre_topc(X45))|(v2_struct_0(X44)|(~v2_pre_topc(X44)|~l1_pre_topc(X44)))))))&((m1_subset_1(esk7_3(X44,X45,X46),k1_zfmisc_1(u1_struct_0(X45)))|~v2_funct_1(X46)|k2_relset_1(u1_struct_0(X44),u1_struct_0(X45),X46)!=k2_struct_0(X45)|k1_relset_1(u1_struct_0(X44),u1_struct_0(X45),X46)!=k2_struct_0(X44)|v3_tops_2(X46,X44,X45)|(~v1_funct_1(X46)|(~v1_funct_2(X46,u1_struct_0(X44),u1_struct_0(X45))|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X44),u1_struct_0(X45))))))|(~v2_pre_topc(X45)|~l1_pre_topc(X45))|(v2_struct_0(X44)|(~v2_pre_topc(X44)|~l1_pre_topc(X44))))&(k8_relset_1(u1_struct_0(X44),u1_struct_0(X45),X46,k2_pre_topc(X45,esk7_3(X44,X45,X46)))!=k2_pre_topc(X44,k8_relset_1(u1_struct_0(X44),u1_struct_0(X45),X46,esk7_3(X44,X45,X46)))|~v2_funct_1(X46)|k2_relset_1(u1_struct_0(X44),u1_struct_0(X45),X46)!=k2_struct_0(X45)|k1_relset_1(u1_struct_0(X44),u1_struct_0(X45),X46)!=k2_struct_0(X44)|v3_tops_2(X46,X44,X45)|(~v1_funct_1(X46)|(~v1_funct_2(X46,u1_struct_0(X44),u1_struct_0(X45))|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X44),u1_struct_0(X45))))))|(~v2_pre_topc(X45)|~l1_pre_topc(X45))|(v2_struct_0(X44)|(~v2_pre_topc(X44)|~l1_pre_topc(X44)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t73_tops_2])])])])])])).
cnf(c_0_82, plain, (v2_struct_0(X2)|v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v3_tops_2(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_83, negated_conjecture, (k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),X1)=k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1)|~l1_struct_0(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_53]), c_0_20])])).
cnf(c_0_84, plain, (m1_subset_1(k2_pre_topc(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_85, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_71])]), c_0_80]), c_0_73]), c_0_74])])).
cnf(c_0_86, plain, (k8_relset_1(u1_struct_0(X3),u1_struct_0(X2),X4,k2_pre_topc(X2,X1))=k2_pre_topc(X3,k8_relset_1(u1_struct_0(X3),u1_struct_0(X2),X4,X1))|v2_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v3_tops_2(X4,X3,X2)|~v1_funct_1(X4)|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_87, negated_conjecture, (v3_tops_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_71]), c_0_18]), c_0_17]), c_0_19]), c_0_20]), c_0_21])]), c_0_24])).
cnf(c_0_88, negated_conjecture, (k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),X1)=k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_53]), c_0_21])])).
cnf(c_0_89, negated_conjecture, (m1_subset_1(k2_pre_topc(esk1_0,esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_21])])).
cnf(c_0_90, negated_conjecture, (k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,k2_pre_topc(esk1_0,esk4_0))!=k2_pre_topc(esk2_0,k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_0))|~v2_funct_1(esk3_0)|k2_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|k1_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_91, negated_conjecture, (k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),k2_pre_topc(esk1_0,X1))=k2_pre_topc(esk2_0,k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_45]), c_0_44]), c_0_46]), c_0_20]), c_0_21]), c_0_22]), c_0_23])]), c_0_24])).
cnf(c_0_92, negated_conjecture, (k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk4_0)=k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_88, c_0_85])).
cnf(c_0_93, negated_conjecture, (k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),k2_pre_topc(esk1_0,esk4_0))=k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,k2_pre_topc(esk1_0,esk4_0))), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_94, negated_conjecture, (k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,k2_pre_topc(esk1_0,esk4_0))!=k2_pre_topc(esk2_0,k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_71])]), c_0_80]), c_0_73]), c_0_74])])).
cnf(c_0_95, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_85]), c_0_92]), c_0_93]), c_0_94]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 96
# Proof object clause steps            : 69
# Proof object formula steps           : 27
# Proof object conjectures             : 50
# Proof object clause conjectures      : 47
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 35
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 31
# Proof object simplifying inferences  : 134
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 46
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 46
# Processed clauses                    : 683
# ...of these trivial                  : 9
# ...subsumed                          : 173
# ...remaining for further processing  : 500
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 16
# Backward-rewritten                   : 264
# Generated clauses                    : 1701
# ...of the previous two non-trivial   : 1238
# Contextual simplify-reflections      : 26
# Paramodulations                      : 1699
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 176
#    Positive orientable unit clauses  : 134
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 39
# Current number of unprocessed clauses: 548
# ...number of literals in the above   : 1102
# Current number of archived formulas  : 0
# Current number of archived clauses   : 322
# Clause-clause subsumption calls (NU) : 19229
# Rec. Clause-clause subsumption calls : 8489
# Non-unit clause-clause subsumptions  : 215
# Unit Clause-clause subsumption calls : 2358
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4906
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 126105

# -------------------------------------------------
# User time                : 0.062 s
# System time              : 0.000 s
# Total time               : 0.062 s
# Maximum resident set size: 28752 pages
