# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/181f468fbe674a1fbf3e6b21c162adbe.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/181f468fbe674a1fbf3e6b21c162adbe.E_input', t2_xboole_1)).
fof(fraenkel_a_2_5_lattice3, axiom, ![X1, X2, X3]:((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))=>(r2_hidden(X1,a_2_5_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r3_lattices(X2,X4,X5)&r2_hidden(X5,X3))))))), file('data/MPTP/proofs/181f468fbe674a1fbf3e6b21c162adbe.E_input', fraenkel_a_2_5_lattice3)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/181f468fbe674a1fbf3e6b21c162adbe.E_input', d3_tarski)).
fof(t3_xboole_1, axiom, ![X1]:(r1_tarski(X1,k1_xboole_0)=>X1=k1_xboole_0), file('data/MPTP/proofs/181f468fbe674a1fbf3e6b21c162adbe.E_input', t3_xboole_1)).
fof(t47_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(k15_lattice3(X1,X2)=k15_lattice3(X1,a_2_5_lattice3(X1,X2))&k16_lattice3(X1,X2)=k16_lattice3(X1,a_2_6_lattice3(X1,X2)))), file('data/MPTP/proofs/181f468fbe674a1fbf3e6b21c162adbe.E_input', t47_lattice3)).
fof(t46_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(r1_tarski(X2,X3)=>(r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3))&r3_lattices(X1,k16_lattice3(X1,X3),k16_lattice3(X1,X2))))), file('data/MPTP/proofs/181f468fbe674a1fbf3e6b21c162adbe.E_input', t46_lattice3)).
fof(t50_lattice3, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))&k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0))), file('data/MPTP/proofs/181f468fbe674a1fbf3e6b21c162adbe.E_input', t50_lattice3)).
fof(t43_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k15_lattice3(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k16_lattice3(X1,k6_domain_1(u1_struct_0(X1),X2))=X2))), file('data/MPTP/proofs/181f468fbe674a1fbf3e6b21c162adbe.E_input', t43_lattice3)).
fof(dt_k15_lattice3, axiom, ![X1, X2]:((~(v2_struct_0(X1))&l3_lattices(X1))=>m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/181f468fbe674a1fbf3e6b21c162adbe.E_input', dt_k15_lattice3)).
fof(t21_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattices(X1,X2,X3)<=>k2_lattices(X1,X2,X3)=X2)))), file('data/MPTP/proofs/181f468fbe674a1fbf3e6b21c162adbe.E_input', t21_lattices)).
fof(redefinition_r3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_lattices(X1,X2,X3)<=>r1_lattices(X1,X2,X3))), file('data/MPTP/proofs/181f468fbe674a1fbf3e6b21c162adbe.E_input', redefinition_r3_lattices)).
fof(d16_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v13_lattices(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(X2=k5_lattices(X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_lattices(X1,X2,X3)=X2&k2_lattices(X1,X3,X2)=X2)))))), file('data/MPTP/proofs/181f468fbe674a1fbf3e6b21c162adbe.E_input', d16_lattices)).
fof(dt_k5_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>m1_subset_1(k5_lattices(X1),u1_struct_0(X1))), file('data/MPTP/proofs/181f468fbe674a1fbf3e6b21c162adbe.E_input', dt_k5_lattices)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/181f468fbe674a1fbf3e6b21c162adbe.E_input', cc1_lattices)).
fof(d6_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v6_lattices(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2))))), file('data/MPTP/proofs/181f468fbe674a1fbf3e6b21c162adbe.E_input', d6_lattices)).
fof(dt_l3_lattices, axiom, ![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1))), file('data/MPTP/proofs/181f468fbe674a1fbf3e6b21c162adbe.E_input', dt_l3_lattices)).
fof(d13_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v13_lattices(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_lattices(X1,X2,X3)=X2&k2_lattices(X1,X3,X2)=X2))))), file('data/MPTP/proofs/181f468fbe674a1fbf3e6b21c162adbe.E_input', d13_lattices)).
fof(c_0_18, plain, ![X6, X7]:(~r2_hidden(X6,X7)|~r1_tarski(X7,X6)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_19, plain, ![X8]:r1_tarski(k1_xboole_0,X8), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_20, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_21, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_22, plain, ![X9, X10, X11, X14, X15, X16]:(((m1_subset_1(esk1_3(X9,X10,X11),u1_struct_0(X10))|~r2_hidden(X9,a_2_5_lattice3(X10,X11))|(v2_struct_0(X10)|(~v10_lattices(X10)|(~v4_lattice3(X10)|~l3_lattices(X10)))))&((X9=esk1_3(X9,X10,X11)|~r2_hidden(X9,a_2_5_lattice3(X10,X11))|(v2_struct_0(X10)|(~v10_lattices(X10)|(~v4_lattice3(X10)|~l3_lattices(X10)))))&((m1_subset_1(esk2_3(X9,X10,X11),u1_struct_0(X10))|~r2_hidden(X9,a_2_5_lattice3(X10,X11))|(v2_struct_0(X10)|(~v10_lattices(X10)|(~v4_lattice3(X10)|~l3_lattices(X10)))))&((r3_lattices(X10,esk1_3(X9,X10,X11),esk2_3(X9,X10,X11))|~r2_hidden(X9,a_2_5_lattice3(X10,X11))|(v2_struct_0(X10)|(~v10_lattices(X10)|(~v4_lattice3(X10)|~l3_lattices(X10)))))&(r2_hidden(esk2_3(X9,X10,X11),X11)|~r2_hidden(X9,a_2_5_lattice3(X10,X11))|(v2_struct_0(X10)|(~v10_lattices(X10)|(~v4_lattice3(X10)|~l3_lattices(X10)))))))))&(~m1_subset_1(X15,u1_struct_0(X10))|(X9!=X15|(~m1_subset_1(X16,u1_struct_0(X10))|(~r3_lattices(X10,X15,X16)|~r2_hidden(X16,X14))))|r2_hidden(X9,a_2_5_lattice3(X10,X14))|(v2_struct_0(X10)|(~v10_lattices(X10)|(~v4_lattice3(X10)|~l3_lattices(X10)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_5_lattice3])])])])])])])).
cnf(c_0_23, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_24, plain, (r2_hidden(esk2_3(X1,X2,X3),X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_5_lattice3(X2,X3))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_25, plain, (v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(X2,a_2_5_lattice3(X1,k1_xboole_0))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_26, plain, ![X17, X18, X19, X20, X21]:((~r1_tarski(X17,X18)|(~r2_hidden(X19,X17)|r2_hidden(X19,X18)))&((r2_hidden(esk3_2(X20,X21),X20)|r1_tarski(X20,X21))&(~r2_hidden(esk3_2(X20,X21),X21)|r1_tarski(X20,X21)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_27, plain, ![X23]:(~r1_tarski(X23,k1_xboole_0)|X23=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_xboole_1])])).
cnf(c_0_28, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~l3_lattices(X2)|~l3_lattices(X1)|~v4_lattice3(X2)|~v4_lattice3(X1)|~v10_lattices(X2)|~v10_lattices(X1)|~r2_hidden(X3,a_2_5_lattice3(X1,a_2_5_lattice3(X2,k1_xboole_0)))), inference(spm,[status(thm)],[c_0_25, c_0_24])).
cnf(c_0_29, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_30, plain, ![X24, X25]:((k15_lattice3(X24,X25)=k15_lattice3(X24,a_2_5_lattice3(X24,X25))|(v2_struct_0(X24)|(~v10_lattices(X24)|(~v4_lattice3(X24)|~l3_lattices(X24)))))&(k16_lattice3(X24,X25)=k16_lattice3(X24,a_2_6_lattice3(X24,X25))|(v2_struct_0(X24)|(~v10_lattices(X24)|(~v4_lattice3(X24)|~l3_lattices(X24)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t47_lattice3])])])])])).
cnf(c_0_31, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_32, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(a_2_5_lattice3(X2,a_2_5_lattice3(X1,k1_xboole_0)),X3)|~l3_lattices(X1)|~l3_lattices(X2)|~v4_lattice3(X1)|~v4_lattice3(X2)|~v10_lattices(X1)|~v10_lattices(X2)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
fof(c_0_33, plain, ![X26, X27, X28]:((r3_lattices(X26,k15_lattice3(X26,X27),k15_lattice3(X26,X28))|~r1_tarski(X27,X28)|(v2_struct_0(X26)|(~v10_lattices(X26)|(~v4_lattice3(X26)|~l3_lattices(X26)))))&(r3_lattices(X26,k16_lattice3(X26,X28),k16_lattice3(X26,X27))|~r1_tarski(X27,X28)|(v2_struct_0(X26)|(~v10_lattices(X26)|(~v4_lattice3(X26)|~l3_lattices(X26)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t46_lattice3])])])])])).
cnf(c_0_34, plain, (k15_lattice3(X1,X2)=k15_lattice3(X1,a_2_5_lattice3(X1,X2))|v2_struct_0(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_35, plain, (a_2_5_lattice3(X1,a_2_5_lattice3(X2,k1_xboole_0))=k1_xboole_0|v2_struct_0(X1)|v2_struct_0(X2)|~l3_lattices(X2)|~l3_lattices(X1)|~v4_lattice3(X2)|~v4_lattice3(X1)|~v10_lattices(X2)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
fof(c_0_36, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))&k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0)))), inference(assume_negation,[status(cth)],[t50_lattice3])).
cnf(c_0_37, plain, (r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3))|v2_struct_0(X1)|~r1_tarski(X2,X3)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_38, plain, (k15_lattice3(X1,a_2_5_lattice3(X2,k1_xboole_0))=k15_lattice3(X1,k1_xboole_0)|v2_struct_0(X2)|v2_struct_0(X1)|~l3_lattices(X1)|~l3_lattices(X2)|~v4_lattice3(X1)|~v4_lattice3(X2)|~v10_lattices(X1)|~v10_lattices(X2)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_39, plain, (v2_struct_0(X1)|r1_tarski(a_2_5_lattice3(X1,k1_xboole_0),X2)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_25, c_0_29])).
fof(c_0_40, negated_conjecture, ((~v2_struct_0(esk4_0)&(v10_lattices(esk4_0)&(v4_lattice3(esk4_0)&l3_lattices(esk4_0))))&(v2_struct_0(esk4_0)|(~v10_lattices(esk4_0)|(~v13_lattices(esk4_0)|~l3_lattices(esk4_0)))|k5_lattices(esk4_0)!=k15_lattice3(esk4_0,k1_xboole_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_36])])])])).
cnf(c_0_41, plain, (r3_lattices(X1,k15_lattice3(X1,k1_xboole_0),k15_lattice3(X1,X2))|v2_struct_0(X3)|v2_struct_0(X1)|~l3_lattices(X1)|~l3_lattices(X3)|~v4_lattice3(X1)|~v4_lattice3(X3)|~v10_lattices(X1)|~v10_lattices(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
cnf(c_0_42, negated_conjecture, (v4_lattice3(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_43, negated_conjecture, (l3_lattices(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_44, negated_conjecture, (v10_lattices(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_45, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_46, plain, ![X30, X31]:((k15_lattice3(X30,k6_domain_1(u1_struct_0(X30),X31))=X31|~m1_subset_1(X31,u1_struct_0(X30))|(v2_struct_0(X30)|(~v10_lattices(X30)|(~v4_lattice3(X30)|~l3_lattices(X30)))))&(k16_lattice3(X30,k6_domain_1(u1_struct_0(X30),X31))=X31|~m1_subset_1(X31,u1_struct_0(X30))|(v2_struct_0(X30)|(~v10_lattices(X30)|(~v4_lattice3(X30)|~l3_lattices(X30)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t43_lattice3])])])])])).
cnf(c_0_47, plain, (r2_hidden(X3,a_2_5_lattice3(X2,X5))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~m1_subset_1(X4,u1_struct_0(X2))|~r3_lattices(X2,X1,X4)|~r2_hidden(X4,X5)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_48, negated_conjecture, (r3_lattices(X1,k15_lattice3(X1,k1_xboole_0),k15_lattice3(X1,X2))|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), c_0_44])]), c_0_45])).
cnf(c_0_49, plain, (k15_lattice3(X1,k6_domain_1(u1_struct_0(X1),X2))=X2|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_50, plain, ![X32, X33]:(v2_struct_0(X32)|~l3_lattices(X32)|m1_subset_1(k15_lattice3(X32,X33),u1_struct_0(X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k15_lattice3])])])).
cnf(c_0_51, plain, (v2_struct_0(X1)|r2_hidden(X2,a_2_5_lattice3(X1,X3))|~r3_lattices(X1,X2,X4)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(X4,X3)), inference(er,[status(thm)],[c_0_47])).
cnf(c_0_52, negated_conjecture, (r3_lattices(X1,k15_lattice3(X1,k1_xboole_0),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_53, plain, (v2_struct_0(X1)|m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_54, negated_conjecture, (v2_struct_0(X1)|r2_hidden(k15_lattice3(X1,k1_xboole_0),a_2_5_lattice3(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(X3,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53])).
cnf(c_0_55, plain, (m1_subset_1(esk2_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_5_lattice3(X2,X3))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_56, negated_conjecture, (v2_struct_0(X1)|r2_hidden(k15_lattice3(X1,k1_xboole_0),a_2_5_lattice3(X1,X2))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(esk2_3(X3,X1,X4),X2)|~r2_hidden(X3,a_2_5_lattice3(X1,X4))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_57, negated_conjecture, (v2_struct_0(X1)|r2_hidden(k15_lattice3(X1,k1_xboole_0),a_2_5_lattice3(X1,X2))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(X3,a_2_5_lattice3(X1,X2))), inference(spm,[status(thm)],[c_0_56, c_0_24])).
cnf(c_0_58, negated_conjecture, (v2_struct_0(X1)|r1_tarski(a_2_5_lattice3(X1,X2),X3)|r2_hidden(k15_lattice3(X1,k1_xboole_0),a_2_5_lattice3(X1,X2))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_57, c_0_29])).
cnf(c_0_59, negated_conjecture, (r1_tarski(a_2_5_lattice3(esk4_0,X1),X2)|r2_hidden(k15_lattice3(esk4_0,k1_xboole_0),a_2_5_lattice3(esk4_0,X1))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_43]), c_0_42]), c_0_44])]), c_0_45])).
cnf(c_0_60, negated_conjecture, (v2_struct_0(X1)|r1_tarski(a_2_5_lattice3(esk4_0,a_2_5_lattice3(X1,k1_xboole_0)),X2)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_59]), c_0_43]), c_0_42]), c_0_44])]), c_0_45])).
cnf(c_0_61, negated_conjecture, (a_2_5_lattice3(esk4_0,a_2_5_lattice3(X1,k1_xboole_0))=k1_xboole_0|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_31, c_0_60])).
fof(c_0_62, plain, ![X34, X35, X36]:((~r1_lattices(X34,X35,X36)|k2_lattices(X34,X35,X36)=X35|~m1_subset_1(X36,u1_struct_0(X34))|~m1_subset_1(X35,u1_struct_0(X34))|(v2_struct_0(X34)|(~v8_lattices(X34)|(~v9_lattices(X34)|~l3_lattices(X34)))))&(k2_lattices(X34,X35,X36)!=X35|r1_lattices(X34,X35,X36)|~m1_subset_1(X36,u1_struct_0(X34))|~m1_subset_1(X35,u1_struct_0(X34))|(v2_struct_0(X34)|(~v8_lattices(X34)|(~v9_lattices(X34)|~l3_lattices(X34)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t21_lattices])])])])])).
fof(c_0_63, plain, ![X37, X38, X39]:((~r3_lattices(X37,X38,X39)|r1_lattices(X37,X38,X39)|(v2_struct_0(X37)|(~v6_lattices(X37)|(~v8_lattices(X37)|(~v9_lattices(X37)|~l3_lattices(X37))))|(~m1_subset_1(X38,u1_struct_0(X37))|~m1_subset_1(X39,u1_struct_0(X37)))))&(~r1_lattices(X37,X38,X39)|r3_lattices(X37,X38,X39)|(v2_struct_0(X37)|(~v6_lattices(X37)|(~v8_lattices(X37)|(~v9_lattices(X37)|~l3_lattices(X37))))|(~m1_subset_1(X38,u1_struct_0(X37))|~m1_subset_1(X39,u1_struct_0(X37)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_lattices])])])])).
cnf(c_0_64, plain, (r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3))|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r1_tarski(a_2_5_lattice3(X1,X2),X3)), inference(spm,[status(thm)],[c_0_37, c_0_34])).
cnf(c_0_65, negated_conjecture, (r1_tarski(a_2_5_lattice3(esk4_0,k1_xboole_0),X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_59]), c_0_43]), c_0_42]), c_0_44])]), c_0_45])).
cnf(c_0_66, negated_conjecture, (k15_lattice3(esk4_0,a_2_5_lattice3(X1,k1_xboole_0))=k15_lattice3(esk4_0,k1_xboole_0)|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_61]), c_0_43]), c_0_42]), c_0_44])]), c_0_45])).
fof(c_0_67, plain, ![X40, X41, X42]:(((k2_lattices(X40,X41,X42)=X41|~m1_subset_1(X42,u1_struct_0(X40))|X41!=k5_lattices(X40)|~m1_subset_1(X41,u1_struct_0(X40))|~v13_lattices(X40)|(v2_struct_0(X40)|~l1_lattices(X40)))&(k2_lattices(X40,X42,X41)=X41|~m1_subset_1(X42,u1_struct_0(X40))|X41!=k5_lattices(X40)|~m1_subset_1(X41,u1_struct_0(X40))|~v13_lattices(X40)|(v2_struct_0(X40)|~l1_lattices(X40))))&((m1_subset_1(esk5_2(X40,X41),u1_struct_0(X40))|X41=k5_lattices(X40)|~m1_subset_1(X41,u1_struct_0(X40))|~v13_lattices(X40)|(v2_struct_0(X40)|~l1_lattices(X40)))&(k2_lattices(X40,X41,esk5_2(X40,X41))!=X41|k2_lattices(X40,esk5_2(X40,X41),X41)!=X41|X41=k5_lattices(X40)|~m1_subset_1(X41,u1_struct_0(X40))|~v13_lattices(X40)|(v2_struct_0(X40)|~l1_lattices(X40))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattices])])])])])])).
fof(c_0_68, plain, ![X44]:(v2_struct_0(X44)|~l1_lattices(X44)|m1_subset_1(k5_lattices(X44),u1_struct_0(X44))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_lattices])])])).
cnf(c_0_69, plain, (k2_lattices(X1,X2,X3)=X2|v2_struct_0(X1)|~r1_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_70, plain, (r1_lattices(X1,X2,X3)|v2_struct_0(X1)|~r3_lattices(X1,X2,X3)|~v6_lattices(X1)|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_71, plain, (r3_lattices(X1,k15_lattice3(X1,X2),X3)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r1_tarski(a_2_5_lattice3(X1,X2),k6_domain_1(u1_struct_0(X1),X3))), inference(spm,[status(thm)],[c_0_64, c_0_49])).
cnf(c_0_72, negated_conjecture, (a_2_5_lattice3(esk4_0,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_31, c_0_65])).
cnf(c_0_73, negated_conjecture, (m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_66]), c_0_43])]), c_0_45])).
cnf(c_0_74, plain, (k2_lattices(X1,X2,X3)=X3|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|X3!=k5_lattices(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v13_lattices(X1)|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_75, plain, (v2_struct_0(X1)|m1_subset_1(k5_lattices(X1),u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_76, plain, (k2_lattices(X1,X2,X3)=X2|v2_struct_0(X1)|~v6_lattices(X1)|~v9_lattices(X1)|~v8_lattices(X1)|~r3_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_77, negated_conjecture, (r3_lattices(esk4_0,k15_lattice3(esk4_0,k1_xboole_0),X1)|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_43]), c_0_42]), c_0_44]), c_0_21])]), c_0_45])).
cnf(c_0_78, negated_conjecture, (m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_42]), c_0_43]), c_0_44])]), c_0_45])).
cnf(c_0_79, negated_conjecture, (v2_struct_0(esk4_0)|~v10_lattices(esk4_0)|~v13_lattices(esk4_0)|~l3_lattices(esk4_0)|k5_lattices(esk4_0)!=k15_lattice3(esk4_0,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_80, plain, (k2_lattices(X1,X2,k5_lattices(X1))=k5_lattices(X1)|v2_struct_0(X1)|~l1_lattices(X1)|~v13_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_74]), c_0_75])).
cnf(c_0_81, negated_conjecture, (k2_lattices(esk4_0,k15_lattice3(esk4_0,k1_xboole_0),X1)=k15_lattice3(esk4_0,k1_xboole_0)|~v6_lattices(esk4_0)|~v9_lattices(esk4_0)|~v8_lattices(esk4_0)|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78]), c_0_43])]), c_0_45])).
cnf(c_0_82, negated_conjecture, (k5_lattices(esk4_0)!=k15_lattice3(esk4_0,k1_xboole_0)|~v13_lattices(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_44]), c_0_43])]), c_0_45])).
cnf(c_0_83, negated_conjecture, (~l1_lattices(esk4_0)|~v6_lattices(esk4_0)|~v9_lattices(esk4_0)|~v8_lattices(esk4_0)|~v13_lattices(esk4_0)|~m1_subset_1(k5_lattices(esk4_0),u1_struct_0(esk4_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_78])]), c_0_45]), c_0_82])).
fof(c_0_84, plain, ![X45]:((~v2_struct_0(X45)|(v2_struct_0(X45)|~v10_lattices(X45))|~l3_lattices(X45))&((v4_lattices(X45)|(v2_struct_0(X45)|~v10_lattices(X45))|~l3_lattices(X45))&((v5_lattices(X45)|(v2_struct_0(X45)|~v10_lattices(X45))|~l3_lattices(X45))&((v6_lattices(X45)|(v2_struct_0(X45)|~v10_lattices(X45))|~l3_lattices(X45))&((v7_lattices(X45)|(v2_struct_0(X45)|~v10_lattices(X45))|~l3_lattices(X45))&((v8_lattices(X45)|(v2_struct_0(X45)|~v10_lattices(X45))|~l3_lattices(X45))&(v9_lattices(X45)|(v2_struct_0(X45)|~v10_lattices(X45))|~l3_lattices(X45)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
fof(c_0_85, plain, ![X46, X47, X48]:((~v6_lattices(X46)|(~m1_subset_1(X47,u1_struct_0(X46))|(~m1_subset_1(X48,u1_struct_0(X46))|k2_lattices(X46,X47,X48)=k2_lattices(X46,X48,X47)))|(v2_struct_0(X46)|~l1_lattices(X46)))&((m1_subset_1(esk6_1(X46),u1_struct_0(X46))|v6_lattices(X46)|(v2_struct_0(X46)|~l1_lattices(X46)))&((m1_subset_1(esk7_1(X46),u1_struct_0(X46))|v6_lattices(X46)|(v2_struct_0(X46)|~l1_lattices(X46)))&(k2_lattices(X46,esk6_1(X46),esk7_1(X46))!=k2_lattices(X46,esk7_1(X46),esk6_1(X46))|v6_lattices(X46)|(v2_struct_0(X46)|~l1_lattices(X46)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_lattices])])])])])])).
fof(c_0_86, plain, ![X51]:((l1_lattices(X51)|~l3_lattices(X51))&(l2_lattices(X51)|~l3_lattices(X51))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
cnf(c_0_87, negated_conjecture, (~l1_lattices(esk4_0)|~v6_lattices(esk4_0)|~v9_lattices(esk4_0)|~v8_lattices(esk4_0)|~v13_lattices(esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_75]), c_0_45])).
cnf(c_0_88, plain, (v9_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_89, plain, (k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2)|v2_struct_0(X1)|~v6_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_90, plain, (v6_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_91, plain, (l1_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_92, negated_conjecture, (~l1_lattices(esk4_0)|~v6_lattices(esk4_0)|~v8_lattices(esk4_0)|~v13_lattices(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_43]), c_0_44])]), c_0_45])).
cnf(c_0_93, plain, (v8_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
fof(c_0_94, plain, ![X52, X54, X55]:(((m1_subset_1(esk8_1(X52),u1_struct_0(X52))|~v13_lattices(X52)|(v2_struct_0(X52)|~l1_lattices(X52)))&((k2_lattices(X52,esk8_1(X52),X54)=esk8_1(X52)|~m1_subset_1(X54,u1_struct_0(X52))|~v13_lattices(X52)|(v2_struct_0(X52)|~l1_lattices(X52)))&(k2_lattices(X52,X54,esk8_1(X52))=esk8_1(X52)|~m1_subset_1(X54,u1_struct_0(X52))|~v13_lattices(X52)|(v2_struct_0(X52)|~l1_lattices(X52)))))&((m1_subset_1(esk9_2(X52,X55),u1_struct_0(X52))|~m1_subset_1(X55,u1_struct_0(X52))|v13_lattices(X52)|(v2_struct_0(X52)|~l1_lattices(X52)))&(k2_lattices(X52,X55,esk9_2(X52,X55))!=X55|k2_lattices(X52,esk9_2(X52,X55),X55)!=X55|~m1_subset_1(X55,u1_struct_0(X52))|v13_lattices(X52)|(v2_struct_0(X52)|~l1_lattices(X52))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_lattices])])])])])])).
cnf(c_0_95, plain, (k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_91])).
cnf(c_0_96, negated_conjecture, (~l1_lattices(esk4_0)|~v6_lattices(esk4_0)|~v13_lattices(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_43]), c_0_44])]), c_0_45])).
cnf(c_0_97, plain, (v13_lattices(X1)|v2_struct_0(X1)|k2_lattices(X1,X2,esk9_2(X1,X2))!=X2|k2_lattices(X1,esk9_2(X1,X2),X2)!=X2|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_98, negated_conjecture, (k2_lattices(esk4_0,X1,k15_lattice3(esk4_0,k1_xboole_0))=k15_lattice3(esk4_0,k1_xboole_0)|~v6_lattices(esk4_0)|~v9_lattices(esk4_0)|~v8_lattices(esk4_0)|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_95]), c_0_78]), c_0_43]), c_0_44])]), c_0_45])).
cnf(c_0_99, negated_conjecture, (~l1_lattices(esk4_0)|~v13_lattices(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_90]), c_0_43]), c_0_44])]), c_0_45])).
cnf(c_0_100, negated_conjecture, (~l1_lattices(esk4_0)|~v6_lattices(esk4_0)|~v9_lattices(esk4_0)|~v8_lattices(esk4_0)|~m1_subset_1(esk9_2(esk4_0,k15_lattice3(esk4_0,k1_xboole_0)),u1_struct_0(esk4_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_78])]), c_0_45]), c_0_81]), c_0_99])).
cnf(c_0_101, plain, (m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|v13_lattices(X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_102, negated_conjecture, (~l1_lattices(esk4_0)|~v6_lattices(esk4_0)|~v9_lattices(esk4_0)|~v8_lattices(esk4_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_78])]), c_0_45]), c_0_99])).
cnf(c_0_103, negated_conjecture, (~l1_lattices(esk4_0)|~v6_lattices(esk4_0)|~v8_lattices(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_88]), c_0_43]), c_0_44])]), c_0_45])).
cnf(c_0_104, negated_conjecture, (~l1_lattices(esk4_0)|~v6_lattices(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_93]), c_0_43]), c_0_44])]), c_0_45])).
cnf(c_0_105, negated_conjecture, (~l1_lattices(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_90]), c_0_43]), c_0_44])]), c_0_45])).
cnf(c_0_106, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_91]), c_0_43])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 107
# Proof object clause steps            : 70
# Proof object formula steps           : 37
# Proof object conjectures             : 37
# Proof object clause conjectures      : 34
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 40
# Proof object simplifying inferences  : 95
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 51
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 50
# Processed clauses                    : 558
# ...of these trivial                  : 1
# ...subsumed                          : 319
# ...remaining for further processing  : 238
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 12
# Backward-rewritten                   : 6
# Generated clauses                    : 1049
# ...of the previous two non-trivial   : 890
# Contextual simplify-reflections      : 25
# Paramodulations                      : 1046
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 167
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 155
# Current number of unprocessed clauses: 427
# ...number of literals in the above   : 2751
# Current number of archived formulas  : 0
# Current number of archived clauses   : 68
# Clause-clause subsumption calls (NU) : 15481
# Rec. Clause-clause subsumption calls : 1532
# Non-unit clause-clause subsumptions  : 265
# Unit Clause-clause subsumption calls : 203
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 33832

# -------------------------------------------------
# User time                : 0.024 s
# System time              : 0.006 s
# Total time               : 0.031 s
# Maximum resident set size: 28876 pages
