# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t2_yellow_0, conjecture, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3, X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(X4=X5=>((r2_lattice3(X1,X3,X4)=>r2_lattice3(X2,X3,X5))&(r1_lattice3(X1,X3,X4)=>r1_lattice3(X2,X3,X5))))))))), file('data/MPTP/proofs/183a599ee1f544b2960fb7a1a5b25917.E_input', t2_yellow_0)).
fof(free_g1_orders_2, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=>![X3, X4]:(g1_orders_2(X1,X2)=g1_orders_2(X3,X4)=>(X1=X3&X2=X4))), file('data/MPTP/proofs/183a599ee1f544b2960fb7a1a5b25917.E_input', free_g1_orders_2)).
fof(dt_u1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))), file('data/MPTP/proofs/183a599ee1f544b2960fb7a1a5b25917.E_input', dt_u1_orders_2)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/183a599ee1f544b2960fb7a1a5b25917.E_input', d8_lattice3)).
fof(t1_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>((X3=X5&X4=X6)=>((r1_orders_2(X1,X3,X4)=>r1_orders_2(X2,X5,X6))&(r2_orders_2(X1,X3,X4)=>r2_orders_2(X2,X5,X6))))))))))), file('data/MPTP/proofs/183a599ee1f544b2960fb7a1a5b25917.E_input', t1_yellow_0)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/183a599ee1f544b2960fb7a1a5b25917.E_input', d9_lattice3)).
fof(c_0_6, negated_conjecture, ~(![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3, X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(X4=X5=>((r2_lattice3(X1,X3,X4)=>r2_lattice3(X2,X3,X5))&(r1_lattice3(X1,X3,X4)=>r1_lattice3(X2,X3,X5)))))))))), inference(assume_negation,[status(cth)],[t2_yellow_0])).
fof(c_0_7, plain, ![X12, X13, X14, X15]:((X12=X14|g1_orders_2(X12,X13)!=g1_orders_2(X14,X15)|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X12,X12))))&(X13=X15|g1_orders_2(X12,X13)!=g1_orders_2(X14,X15)|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X12,X12))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_orders_2])])])])).
fof(c_0_8, negated_conjecture, (l1_orders_2(esk1_0)&(l1_orders_2(esk2_0)&(g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))&(m1_subset_1(esk4_0,u1_struct_0(esk1_0))&(m1_subset_1(esk5_0,u1_struct_0(esk2_0))&(esk4_0=esk5_0&(((r1_lattice3(esk1_0,esk3_0,esk4_0)|r2_lattice3(esk1_0,esk3_0,esk4_0))&(~r1_lattice3(esk2_0,esk3_0,esk5_0)|r2_lattice3(esk1_0,esk3_0,esk4_0)))&((r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_lattice3(esk2_0,esk3_0,esk5_0))&(~r1_lattice3(esk2_0,esk3_0,esk5_0)|~r2_lattice3(esk2_0,esk3_0,esk5_0)))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
cnf(c_0_9, plain, (X1=X2|g1_orders_2(X1,X3)!=g1_orders_2(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10, negated_conjecture, (g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_11, plain, ![X16]:(~l1_orders_2(X16)|m1_subset_1(u1_orders_2(X16),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X16),u1_struct_0(X16))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_orders_2])])).
cnf(c_0_12, negated_conjecture, (X1=u1_struct_0(esk2_0)|g1_orders_2(X1,X2)!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_13, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_14, plain, ![X17, X18, X19, X20]:((~r1_lattice3(X17,X18,X19)|(~m1_subset_1(X20,u1_struct_0(X17))|(~r2_hidden(X20,X18)|r1_orders_2(X17,X19,X20)))|~m1_subset_1(X19,u1_struct_0(X17))|~l1_orders_2(X17))&((m1_subset_1(esk6_3(X17,X18,X19),u1_struct_0(X17))|r1_lattice3(X17,X18,X19)|~m1_subset_1(X19,u1_struct_0(X17))|~l1_orders_2(X17))&((r2_hidden(esk6_3(X17,X18,X19),X18)|r1_lattice3(X17,X18,X19)|~m1_subset_1(X19,u1_struct_0(X17))|~l1_orders_2(X17))&(~r1_orders_2(X17,X19,esk6_3(X17,X18,X19))|r1_lattice3(X17,X18,X19)|~m1_subset_1(X19,u1_struct_0(X17))|~l1_orders_2(X17))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
cnf(c_0_15, negated_conjecture, (u1_struct_0(X1)=u1_struct_0(esk2_0)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_16, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_17, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, negated_conjecture, (u1_struct_0(esk2_0)=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_15]), c_0_16])])).
cnf(c_0_19, negated_conjecture, (l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_20, plain, ![X22, X23, X24, X25, X26, X27]:((~r1_orders_2(X22,X24,X25)|r1_orders_2(X23,X26,X27)|(X24!=X26|X25!=X27)|~m1_subset_1(X27,u1_struct_0(X23))|~m1_subset_1(X26,u1_struct_0(X23))|~m1_subset_1(X25,u1_struct_0(X22))|~m1_subset_1(X24,u1_struct_0(X22))|g1_orders_2(u1_struct_0(X22),u1_orders_2(X22))!=g1_orders_2(u1_struct_0(X23),u1_orders_2(X23))|~l1_orders_2(X23)|~l1_orders_2(X22))&(~r2_orders_2(X22,X24,X25)|r2_orders_2(X23,X26,X27)|(X24!=X26|X25!=X27)|~m1_subset_1(X27,u1_struct_0(X23))|~m1_subset_1(X26,u1_struct_0(X23))|~m1_subset_1(X25,u1_struct_0(X22))|~m1_subset_1(X24,u1_struct_0(X22))|g1_orders_2(u1_struct_0(X22),u1_orders_2(X22))!=g1_orders_2(u1_struct_0(X23),u1_orders_2(X23))|~l1_orders_2(X23)|~l1_orders_2(X22))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_yellow_0])])])])).
cnf(c_0_21, negated_conjecture, (r1_lattice3(esk2_0,X1,X2)|m1_subset_1(esk6_3(esk2_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])])).
cnf(c_0_22, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_23, plain, ![X28, X29, X30, X31]:((~r2_lattice3(X28,X29,X30)|(~m1_subset_1(X31,u1_struct_0(X28))|(~r2_hidden(X31,X29)|r1_orders_2(X28,X31,X30)))|~m1_subset_1(X30,u1_struct_0(X28))|~l1_orders_2(X28))&((m1_subset_1(esk7_3(X28,X29,X30),u1_struct_0(X28))|r2_lattice3(X28,X29,X30)|~m1_subset_1(X30,u1_struct_0(X28))|~l1_orders_2(X28))&((r2_hidden(esk7_3(X28,X29,X30),X29)|r2_lattice3(X28,X29,X30)|~m1_subset_1(X30,u1_struct_0(X28))|~l1_orders_2(X28))&(~r1_orders_2(X28,esk7_3(X28,X29,X30),X30)|r2_lattice3(X28,X29,X30)|~m1_subset_1(X30,u1_struct_0(X28))|~l1_orders_2(X28))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
cnf(c_0_24, plain, (r1_orders_2(X4,X5,X6)|~r1_orders_2(X1,X2,X3)|X2!=X5|X3!=X6|~m1_subset_1(X6,u1_struct_0(X4))|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))|~l1_orders_2(X4)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_26, negated_conjecture, (r1_lattice3(esk2_0,X1,esk4_0)|m1_subset_1(esk6_3(esk2_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28, plain, (r1_orders_2(X1,X2,X3)|g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_orders_2(X4,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))|~l1_orders_2(X1)|~l1_orders_2(X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_24])])).
cnf(c_0_29, negated_conjecture, (g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk2_0))=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))), inference(rw,[status(thm)],[c_0_10, c_0_18])).
cnf(c_0_30, negated_conjecture, (r1_orders_2(esk1_0,X1,esk6_3(esk2_0,X2,esk4_0))|r1_lattice3(esk2_0,X2,esk4_0)|~r2_hidden(esk6_3(esk2_0,X2,esk4_0),X3)|~r1_lattice3(esk1_0,X3,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_16])])).
cnf(c_0_31, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_32, negated_conjecture, (r2_lattice3(esk2_0,X1,X2)|m1_subset_1(esk7_3(esk2_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_18]), c_0_19])])).
cnf(c_0_33, negated_conjecture, (r1_orders_2(esk2_0,X1,X2)|g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~r1_orders_2(X3,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_18]), c_0_29]), c_0_19])])).
cnf(c_0_34, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,esk6_3(esk2_0,X1,esk4_0))|r1_lattice3(esk2_0,X1,esk4_0)|~r2_hidden(esk6_3(esk2_0,X1,esk4_0),X2)|~r1_lattice3(esk1_0,X2,esk4_0)), inference(spm,[status(thm)],[c_0_30, c_0_22])).
cnf(c_0_35, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk4_0)|r2_lattice3(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_36, negated_conjecture, (r2_hidden(esk6_3(esk2_0,X1,X2),X1)|r1_lattice3(esk2_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_18]), c_0_19])])).
cnf(c_0_37, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk2_0,esk3_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_38, negated_conjecture, (esk4_0=esk5_0), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_39, plain, (r1_orders_2(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_40, negated_conjecture, (r2_lattice3(esk2_0,X1,esk4_0)|m1_subset_1(esk7_3(esk2_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_32, c_0_22])).
cnf(c_0_41, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk6_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_42, negated_conjecture, (r1_orders_2(esk2_0,X1,X2)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_33]), c_0_16])])).
cnf(c_0_43, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,esk6_3(esk2_0,X1,esk4_0))|r1_lattice3(esk2_0,X1,esk4_0)|r2_lattice3(esk1_0,esk3_0,esk4_0)|~r2_hidden(esk6_3(esk2_0,X1,esk4_0),esk3_0)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_44, negated_conjecture, (r2_hidden(esk6_3(esk2_0,X1,esk4_0),X1)|r1_lattice3(esk2_0,X1,esk4_0)), inference(spm,[status(thm)],[c_0_36, c_0_22])).
cnf(c_0_45, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk2_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_46, negated_conjecture, (r1_orders_2(esk1_0,esk7_3(esk2_0,X1,esk4_0),X2)|r2_lattice3(esk2_0,X1,esk4_0)|~r2_hidden(esk7_3(esk2_0,X1,esk4_0),X3)|~r2_lattice3(esk1_0,X3,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_16])])).
cnf(c_0_47, negated_conjecture, (r1_lattice3(esk2_0,X1,X2)|~r1_orders_2(esk1_0,X2,esk6_3(esk2_0,X1,X2))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_18]), c_0_19])]), c_0_21])).
cnf(c_0_48, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,esk6_3(esk2_0,esk3_0,esk4_0))|r2_lattice3(esk1_0,esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])).
cnf(c_0_49, plain, (r2_hidden(esk7_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_50, negated_conjecture, (r1_orders_2(esk1_0,esk7_3(esk2_0,X1,esk4_0),esk4_0)|r2_lattice3(esk2_0,X1,esk4_0)|~r2_hidden(esk7_3(esk2_0,X1,esk4_0),X2)|~r2_lattice3(esk1_0,X2,esk4_0)), inference(spm,[status(thm)],[c_0_46, c_0_22])).
cnf(c_0_51, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_22])]), c_0_45])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk7_3(esk2_0,X1,X2),X1)|r2_lattice3(esk2_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_18]), c_0_19])])).
cnf(c_0_53, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_lattice3(esk2_0,esk3_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_54, plain, (r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk7_3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_55, negated_conjecture, (r1_orders_2(esk1_0,esk7_3(esk2_0,X1,esk4_0),esk4_0)|r2_lattice3(esk2_0,X1,esk4_0)|~r2_hidden(esk7_3(esk2_0,X1,esk4_0),esk3_0)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_56, negated_conjecture, (r2_hidden(esk7_3(esk2_0,X1,esk4_0),X1)|r2_lattice3(esk2_0,X1,esk4_0)), inference(spm,[status(thm)],[c_0_52, c_0_22])).
cnf(c_0_57, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_lattice3(esk2_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_53, c_0_38])).
cnf(c_0_58, negated_conjecture, (r2_lattice3(esk2_0,X1,X2)|~r1_orders_2(esk1_0,esk7_3(esk2_0,X1,X2),X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_42]), c_0_18]), c_0_19])]), c_0_32])).
cnf(c_0_59, negated_conjecture, (r1_orders_2(esk1_0,esk7_3(esk2_0,esk3_0,esk4_0),esk4_0)|r2_lattice3(esk2_0,esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_60, negated_conjecture, (~r1_lattice3(esk2_0,esk3_0,esk5_0)|~r2_lattice3(esk2_0,esk3_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_61, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,esk6_3(esk2_0,X1,esk4_0))|r1_lattice3(esk2_0,X1,esk4_0)|~r2_hidden(esk6_3(esk2_0,X1,esk4_0),esk3_0)|~r2_lattice3(esk2_0,esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_34, c_0_57])).
cnf(c_0_62, negated_conjecture, (r2_lattice3(esk2_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_22])])).
cnf(c_0_63, negated_conjecture, (~r1_lattice3(esk2_0,esk3_0,esk4_0)|~r2_lattice3(esk2_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_38]), c_0_38])).
cnf(c_0_64, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,esk6_3(esk2_0,X1,esk4_0))|r1_lattice3(esk2_0,X1,esk4_0)|~r2_hidden(esk6_3(esk2_0,X1,esk4_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_62])])).
cnf(c_0_65, negated_conjecture, (~r1_lattice3(esk2_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_62])])).
cnf(c_0_66, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,esk6_3(esk2_0,esk3_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_44]), c_0_65])).
cnf(c_0_67, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_66]), c_0_22])]), c_0_65]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 68
# Proof object clause steps            : 55
# Proof object formula steps           : 13
# Proof object conjectures             : 46
# Proof object clause conjectures      : 43
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 28
# Proof object simplifying inferences  : 48
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 128
# ...of these trivial                  : 1
# ...subsumed                          : 2
# ...remaining for further processing  : 125
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 23
# Generated clauses                    : 275
# ...of the previous two non-trivial   : 270
# Contextual simplify-reflections      : 4
# Paramodulations                      : 262
# Factorizations                       : 0
# Equation resolutions                 : 15
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 77
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 65
# Current number of unprocessed clauses: 178
# ...number of literals in the above   : 825
# Current number of archived formulas  : 0
# Current number of archived clauses   : 46
# Clause-clause subsumption calls (NU) : 1439
# Rec. Clause-clause subsumption calls : 284
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 10099

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.003 s
# Total time               : 0.014 s
# Maximum resident set size: 28576 pages
