# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', d4_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', t28_xboole_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', t2_xboole_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', d5_xboole_0)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', t100_xboole_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', d1_xboole_0)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', d10_xboole_0)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', t3_xboole_0)).
fof(t48_zfmisc_1, conjecture, ![X1, X2, X3]:((r2_hidden(X1,X2)&r2_hidden(X3,X2))=>k2_xboole_0(k2_tarski(X1,X3),X2)=X2), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', t48_zfmisc_1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', t73_enumset1)).
fof(t98_enumset1, axiom, ![X1, X2, X3]:k1_enumset1(X1,X2,X3)=k1_enumset1(X1,X3,X2), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', t98_enumset1)).
fof(t137_enumset1, axiom, ![X1, X2, X3]:k2_xboole_0(k2_tarski(X2,X1),k2_tarski(X3,X1))=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', t137_enumset1)).
fof(l53_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X4)), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', l53_enumset1)).
fof(t46_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X4)), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', t46_enumset1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', t69_enumset1)).
fof(t38_zfmisc_1, axiom, ![X1, X2, X3]:(r1_tarski(k2_tarski(X1,X2),X3)<=>(r2_hidden(X1,X3)&r2_hidden(X2,X3))), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', t38_zfmisc_1)).
fof(t2_tarski, axiom, ![X1, X2]:(![X3]:(r2_hidden(X3,X1)<=>r2_hidden(X3,X2))=>X1=X2), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', t2_tarski)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', t39_xboole_1)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/18489427bf924268902c24c0a3df91e2.E_input', t1_boole)).
fof(c_0_23, plain, ![X56, X57, X58, X59, X60, X61, X62, X63]:((((r2_hidden(X59,X56)|~r2_hidden(X59,X58)|X58!=k3_xboole_0(X56,X57))&(r2_hidden(X59,X57)|~r2_hidden(X59,X58)|X58!=k3_xboole_0(X56,X57)))&(~r2_hidden(X60,X56)|~r2_hidden(X60,X57)|r2_hidden(X60,X58)|X58!=k3_xboole_0(X56,X57)))&((~r2_hidden(esk6_3(X61,X62,X63),X63)|(~r2_hidden(esk6_3(X61,X62,X63),X61)|~r2_hidden(esk6_3(X61,X62,X63),X62))|X63=k3_xboole_0(X61,X62))&((r2_hidden(esk6_3(X61,X62,X63),X61)|r2_hidden(esk6_3(X61,X62,X63),X63)|X63=k3_xboole_0(X61,X62))&(r2_hidden(esk6_3(X61,X62,X63),X62)|r2_hidden(esk6_3(X61,X62,X63),X63)|X63=k3_xboole_0(X61,X62))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_24, plain, ![X26, X27]:(~r1_tarski(X26,X27)|k3_xboole_0(X26,X27)=X26), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_25, plain, ![X77]:r1_tarski(k1_xboole_0,X77), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_26, plain, ![X28, X29, X30, X31, X32, X33, X34, X35]:((((r2_hidden(X31,X28)|~r2_hidden(X31,X30)|X30!=k4_xboole_0(X28,X29))&(~r2_hidden(X31,X29)|~r2_hidden(X31,X30)|X30!=k4_xboole_0(X28,X29)))&(~r2_hidden(X32,X28)|r2_hidden(X32,X29)|r2_hidden(X32,X30)|X30!=k4_xboole_0(X28,X29)))&((~r2_hidden(esk4_3(X33,X34,X35),X35)|(~r2_hidden(esk4_3(X33,X34,X35),X33)|r2_hidden(esk4_3(X33,X34,X35),X34))|X35=k4_xboole_0(X33,X34))&((r2_hidden(esk4_3(X33,X34,X35),X33)|r2_hidden(esk4_3(X33,X34,X35),X35)|X35=k4_xboole_0(X33,X34))&(~r2_hidden(esk4_3(X33,X34,X35),X34)|r2_hidden(esk4_3(X33,X34,X35),X35)|X35=k4_xboole_0(X33,X34))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_27, plain, ![X37, X38]:k4_xboole_0(X37,X38)=k5_xboole_0(X37,k3_xboole_0(X37,X38)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_28, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_28])).
cnf(c_0_34, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_29, c_0_30])).
fof(c_0_35, plain, ![X71, X72, X73]:((~v1_xboole_0(X71)|~r2_hidden(X72,X71))&(r2_hidden(esk8_1(X73),X73)|v1_xboole_0(X73))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_36, plain, ![X13, X14]:(((r1_tarski(X13,X14)|X13!=X14)&(r1_tarski(X14,X13)|X13!=X14))&(~r1_tarski(X13,X14)|~r1_tarski(X14,X13)|X13=X14)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_37, plain, ![X65, X66, X68, X69, X70]:(((r2_hidden(esk7_2(X65,X66),X65)|r1_xboole_0(X65,X66))&(r2_hidden(esk7_2(X65,X66),X66)|r1_xboole_0(X65,X66)))&(~r2_hidden(X70,X68)|~r2_hidden(X70,X69)|~r1_xboole_0(X68,X69))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_38, plain, (X3!=k5_xboole_0(X4,k3_xboole_0(X4,X2))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_39, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_40, plain, (r2_hidden(esk8_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_41, negated_conjecture, ~(![X1, X2, X3]:((r2_hidden(X1,X2)&r2_hidden(X3,X2))=>k2_xboole_0(k2_tarski(X1,X3),X2)=X2)), inference(assume_negation,[status(cth)],[t48_zfmisc_1])).
fof(c_0_42, plain, ![X16, X17]:k3_tarski(k2_tarski(X16,X17))=k2_xboole_0(X16,X17), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_43, plain, ![X75, X76]:k1_enumset1(X75,X75,X76)=k2_tarski(X75,X76), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_44, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_45, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_46, plain, (r2_hidden(esk7_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_47, plain, (~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_48, plain, (v1_xboole_0(k1_xboole_0)|r2_hidden(esk8_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
fof(c_0_49, negated_conjecture, ((r2_hidden(esk1_0,esk2_0)&r2_hidden(esk3_0,esk2_0))&k2_xboole_0(k2_tarski(esk1_0,esk3_0),esk2_0)!=esk2_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_41])])])).
cnf(c_0_50, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_51, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_52, plain, ![X23, X24, X25]:k2_enumset1(X23,X23,X24,X25)=k1_enumset1(X23,X24,X25), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_53, plain, ![X9, X10, X11, X12]:k3_enumset1(X9,X9,X10,X11,X12)=k2_enumset1(X9,X10,X11,X12), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_54, plain, ![X18, X19, X20, X21, X22]:k4_enumset1(X18,X18,X19,X20,X21,X22)=k3_enumset1(X18,X19,X20,X21,X22), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_55, plain, ![X43, X44, X45]:k1_enumset1(X43,X44,X45)=k1_enumset1(X43,X45,X44), inference(variable_rename,[status(thm)],[t98_enumset1])).
fof(c_0_56, plain, ![X81, X82, X83]:k2_xboole_0(k2_tarski(X82,X81),k2_tarski(X83,X81))=k1_enumset1(X81,X82,X83), inference(variable_rename,[status(thm)],[t137_enumset1])).
fof(c_0_57, plain, ![X50, X51, X52, X53]:k2_enumset1(X50,X51,X52,X53)=k2_xboole_0(k2_tarski(X50,X51),k2_tarski(X52,X53)), inference(variable_rename,[status(thm)],[l53_enumset1])).
cnf(c_0_58, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_59, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_60, plain, (r1_xboole_0(X1,X2)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_61, plain, (v1_xboole_0(k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_48])).
cnf(c_0_62, negated_conjecture, (k2_xboole_0(k2_tarski(esk1_0,esk3_0),esk2_0)!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_63, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_64, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_65, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_66, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_67, plain, (k1_enumset1(X1,X2,X3)=k1_enumset1(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_68, plain, (k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X2))=k1_enumset1(X2,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_69, plain, (k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X4))), inference(split_conjunct,[status(thm)],[c_0_57])).
fof(c_0_70, plain, ![X39, X40, X41, X42]:k2_enumset1(X39,X40,X41,X42)=k2_xboole_0(k1_enumset1(X39,X40,X41),k1_tarski(X42)), inference(variable_rename,[status(thm)],[t46_enumset1])).
fof(c_0_71, plain, ![X49]:k2_tarski(X49,X49)=k1_tarski(X49), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_72, plain, ![X78, X79, X80]:(((r2_hidden(X78,X80)|~r1_tarski(k2_tarski(X78,X79),X80))&(r2_hidden(X79,X80)|~r1_tarski(k2_tarski(X78,X79),X80)))&(~r2_hidden(X78,X80)|~r2_hidden(X79,X80)|r1_tarski(k2_tarski(X78,X79),X80))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t38_zfmisc_1])])])).
cnf(c_0_73, plain, (r2_hidden(X1,X2)|X3!=k5_xboole_0(X2,k3_xboole_0(X2,X4))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_58, c_0_32])).
cnf(c_0_74, plain, (k3_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_29, c_0_59])).
cnf(c_0_75, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_76, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
fof(c_0_77, plain, ![X46, X47]:((~r2_hidden(esk5_2(X46,X47),X46)|~r2_hidden(esk5_2(X46,X47),X47)|X46=X47)&(r2_hidden(esk5_2(X46,X47),X46)|r2_hidden(esk5_2(X46,X47),X47)|X46=X47)), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tarski])])])])).
cnf(c_0_78, negated_conjecture, (k3_tarski(k4_enumset1(k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0),k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0),k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0),k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0),k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0),esk2_0))!=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_51]), c_0_63]), c_0_64]), c_0_64]), c_0_64]), c_0_65]), c_0_65]), c_0_65]), c_0_65]), c_0_66]), c_0_66]), c_0_66]), c_0_66]), c_0_66])).
cnf(c_0_79, plain, (k4_enumset1(X1,X1,X1,X1,X2,X3)=k4_enumset1(X1,X1,X1,X1,X3,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_64]), c_0_64]), c_0_65]), c_0_65]), c_0_66]), c_0_66])).
cnf(c_0_80, plain, (k3_tarski(k4_enumset1(k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X3,X3,X3,X3,X3,X2)))=k4_enumset1(X2,X2,X2,X2,X1,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_51]), c_0_51]), c_0_63]), c_0_64]), c_0_64]), c_0_64]), c_0_64]), c_0_64]), c_0_65]), c_0_65]), c_0_65]), c_0_65]), c_0_65]), c_0_65]), c_0_66]), c_0_66]), c_0_66]), c_0_66]), c_0_66]), c_0_66]), c_0_66])).
cnf(c_0_81, plain, (k4_enumset1(X1,X1,X1,X2,X3,X4)=k3_tarski(k4_enumset1(k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X3,X3,X3,X3,X3,X4)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_51]), c_0_51]), c_0_63]), c_0_64]), c_0_64]), c_0_64]), c_0_64]), c_0_65]), c_0_65]), c_0_65]), c_0_65]), c_0_65]), c_0_65]), c_0_66]), c_0_66]), c_0_66]), c_0_66]), c_0_66]), c_0_66]), c_0_66])).
cnf(c_0_82, plain, (k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X4))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_83, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
fof(c_0_84, plain, ![X54, X55]:k2_xboole_0(X54,k4_xboole_0(X55,X54))=k2_xboole_0(X54,X55), inference(variable_rename,[status(thm)],[t39_xboole_1])).
cnf(c_0_85, plain, (r1_tarski(k2_tarski(X1,X3),X2)|~r2_hidden(X1,X2)|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_86, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))), inference(er,[status(thm)],[c_0_73])).
cnf(c_0_87, plain, (~r2_hidden(X1,k5_xboole_0(X2,X2))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_47, c_0_74])).
cnf(c_0_88, plain, (~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_39])).
cnf(c_0_89, plain, (r2_hidden(esk5_2(X1,X2),X1)|r2_hidden(esk5_2(X1,X2),X2)|X1=X2), inference(split_conjunct,[status(thm)],[c_0_77])).
fof(c_0_90, plain, ![X15]:k2_xboole_0(X15,k1_xboole_0)=X15, inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_91, negated_conjecture, (k3_tarski(k4_enumset1(k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0),k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0),k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0),k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0),esk2_0,k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0)))!=esk2_0), inference(rw,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_92, plain, (k4_enumset1(X1,X1,X1,X2,X3,X2)=k4_enumset1(X2,X2,X2,X2,X1,X3)), inference(rw,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_93, plain, (k4_enumset1(X1,X1,X1,X2,X3,X4)=k3_tarski(k4_enumset1(k4_enumset1(X1,X1,X1,X1,X2,X3),k4_enumset1(X1,X1,X1,X1,X2,X3),k4_enumset1(X1,X1,X1,X1,X2,X3),k4_enumset1(X1,X1,X1,X1,X2,X3),k4_enumset1(X1,X1,X1,X1,X2,X3),k4_enumset1(X4,X4,X4,X4,X4,X4)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_83]), c_0_51]), c_0_63]), c_0_64]), c_0_64]), c_0_64]), c_0_64]), c_0_65]), c_0_65]), c_0_65]), c_0_65]), c_0_65]), c_0_65]), c_0_66]), c_0_66]), c_0_66]), c_0_66]), c_0_66]), c_0_66]), c_0_66])).
cnf(c_0_94, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_95, plain, (r1_tarski(k4_enumset1(X1,X1,X1,X1,X1,X3),X2)|~r2_hidden(X3,X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_51]), c_0_64]), c_0_65]), c_0_66])).
cnf(c_0_96, plain, (~r2_hidden(X1,k5_xboole_0(X2,X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_74]), c_0_87])).
cnf(c_0_97, plain, (k1_xboole_0=X1|r2_hidden(esk5_2(k1_xboole_0,X1),X1)), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_98, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_99, negated_conjecture, (k3_tarski(k4_enumset1(esk2_0,esk2_0,esk2_0,k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0),k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0),k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0)))!=esk2_0), inference(rw,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_100, plain, (k4_enumset1(X1,X1,X1,X2,X3,X3)=k4_enumset1(X1,X1,X1,X1,X2,X3)), inference(spm,[status(thm)],[c_0_93, c_0_81])).
cnf(c_0_101, plain, (k3_tarski(k4_enumset1(X1,X1,X1,X1,X1,k5_xboole_0(X2,k3_xboole_0(X2,X1))))=k3_tarski(k4_enumset1(X1,X1,X1,X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_94, c_0_63]), c_0_63]), c_0_32]), c_0_64]), c_0_64]), c_0_65]), c_0_65]), c_0_66]), c_0_66])).
cnf(c_0_102, plain, (k3_xboole_0(k4_enumset1(X1,X1,X1,X1,X1,X2),X3)=k4_enumset1(X1,X1,X1,X1,X1,X2)|~r2_hidden(X2,X3)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_29, c_0_95])).
cnf(c_0_103, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_104, plain, (k3_tarski(k4_enumset1(X1,X1,X1,X1,X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_98, c_0_63]), c_0_64]), c_0_65]), c_0_66])).
cnf(c_0_105, negated_conjecture, (k3_tarski(k4_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0)))!=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_99, c_0_100]), c_0_100])).
cnf(c_0_106, plain, (k3_tarski(k4_enumset1(X1,X1,X1,X1,X1,k4_enumset1(X2,X2,X2,X2,X2,X3)))=X1|~r2_hidden(X3,X1)|~r2_hidden(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_103]), c_0_104])).
cnf(c_0_107, negated_conjecture, (r2_hidden(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_108, negated_conjecture, (r2_hidden(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_109, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_107]), c_0_108])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 110
# Proof object clause steps            : 63
# Proof object formula steps           : 47
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 23
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 118
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 23
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 43
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 36
# Processed clauses                    : 10896
# ...of these trivial                  : 177
# ...subsumed                          : 9707
# ...remaining for further processing  : 1012
# Other redundant clauses eliminated   : 10
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 29
# Backward-rewritten                   : 39
# Generated clauses                    : 143236
# ...of the previous two non-trivial   : 121407
# Contextual simplify-reflections      : 19
# Paramodulations                      : 143080
# Factorizations                       : 146
# Equation resolutions                 : 10
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 901
#    Positive orientable unit clauses  : 101
#    Positive unorientable unit clauses: 26
#    Negative unit clauses             : 10
#    Non-unit-clauses                  : 764
# Current number of unprocessed clauses: 109880
# ...number of literals in the above   : 406041
# Current number of archived formulas  : 0
# Current number of archived clauses   : 110
# Clause-clause subsumption calls (NU) : 562094
# Rec. Clause-clause subsumption calls : 327873
# Non-unit clause-clause subsumptions  : 5846
# Unit Clause-clause subsumption calls : 5439
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5761
# BW rewrite match successes           : 1912
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2131203

# -------------------------------------------------
# User time                : 1.233 s
# System time              : 0.043 s
# Total time               : 1.277 s
# Maximum resident set size: 83024 pages
