# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d8_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(X3=k5_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>?[X6]:(r2_hidden(k4_tarski(X4,X6),X1)&r2_hidden(k4_tarski(X6,X5),X2))))))), file('data/MPTP/proofs/18e00dab055f4d10bf367612e3f8b7ff.E_input', d8_relat_1)).
fof(dt_k5_relat_1, axiom, ![X1, X2]:((v1_relat_1(X1)&v1_relat_1(X2))=>v1_relat_1(k5_relat_1(X1,X2))), file('data/MPTP/proofs/18e00dab055f4d10bf367612e3f8b7ff.E_input', dt_k5_relat_1)).
fof(t8_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>(r2_hidden(X1,k1_relat_1(X3))&X2=k1_funct_1(X3,X1)))), file('data/MPTP/proofs/18e00dab055f4d10bf367612e3f8b7ff.E_input', t8_funct_1)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/18e00dab055f4d10bf367612e3f8b7ff.E_input', d4_relat_1)).
fof(t27_funct_1, conjecture, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(k1_relat_1(k5_relat_1(X2,X1))=k1_relat_1(X2)=>r1_tarski(k2_relat_1(X2),k1_relat_1(X1))))), file('data/MPTP/proofs/18e00dab055f4d10bf367612e3f8b7ff.E_input', t27_funct_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/18e00dab055f4d10bf367612e3f8b7ff.E_input', d5_relat_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/18e00dab055f4d10bf367612e3f8b7ff.E_input', d3_tarski)).
fof(c_0_7, plain, ![X7, X8, X9, X10, X11, X13, X14, X15, X18]:((((r2_hidden(k4_tarski(X10,esk1_5(X7,X8,X9,X10,X11)),X7)|~r2_hidden(k4_tarski(X10,X11),X9)|X9!=k5_relat_1(X7,X8)|~v1_relat_1(X9)|~v1_relat_1(X8)|~v1_relat_1(X7))&(r2_hidden(k4_tarski(esk1_5(X7,X8,X9,X10,X11),X11),X8)|~r2_hidden(k4_tarski(X10,X11),X9)|X9!=k5_relat_1(X7,X8)|~v1_relat_1(X9)|~v1_relat_1(X8)|~v1_relat_1(X7)))&(~r2_hidden(k4_tarski(X13,X15),X7)|~r2_hidden(k4_tarski(X15,X14),X8)|r2_hidden(k4_tarski(X13,X14),X9)|X9!=k5_relat_1(X7,X8)|~v1_relat_1(X9)|~v1_relat_1(X8)|~v1_relat_1(X7)))&((~r2_hidden(k4_tarski(esk2_3(X7,X8,X9),esk3_3(X7,X8,X9)),X9)|(~r2_hidden(k4_tarski(esk2_3(X7,X8,X9),X18),X7)|~r2_hidden(k4_tarski(X18,esk3_3(X7,X8,X9)),X8))|X9=k5_relat_1(X7,X8)|~v1_relat_1(X9)|~v1_relat_1(X8)|~v1_relat_1(X7))&((r2_hidden(k4_tarski(esk2_3(X7,X8,X9),esk4_3(X7,X8,X9)),X7)|r2_hidden(k4_tarski(esk2_3(X7,X8,X9),esk3_3(X7,X8,X9)),X9)|X9=k5_relat_1(X7,X8)|~v1_relat_1(X9)|~v1_relat_1(X8)|~v1_relat_1(X7))&(r2_hidden(k4_tarski(esk4_3(X7,X8,X9),esk3_3(X7,X8,X9)),X8)|r2_hidden(k4_tarski(esk2_3(X7,X8,X9),esk3_3(X7,X8,X9)),X9)|X9=k5_relat_1(X7,X8)|~v1_relat_1(X9)|~v1_relat_1(X8)|~v1_relat_1(X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_relat_1])])])])])])).
fof(c_0_8, plain, ![X20, X21]:(~v1_relat_1(X20)|~v1_relat_1(X21)|v1_relat_1(k5_relat_1(X20,X21))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_relat_1])])).
fof(c_0_9, plain, ![X22, X23, X24]:(((r2_hidden(X22,k1_relat_1(X24))|~r2_hidden(k4_tarski(X22,X23),X24)|(~v1_relat_1(X24)|~v1_funct_1(X24)))&(X23=k1_funct_1(X24,X22)|~r2_hidden(k4_tarski(X22,X23),X24)|(~v1_relat_1(X24)|~v1_funct_1(X24))))&(~r2_hidden(X22,k1_relat_1(X24))|X23!=k1_funct_1(X24,X22)|r2_hidden(k4_tarski(X22,X23),X24)|(~v1_relat_1(X24)|~v1_funct_1(X24)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_funct_1])])])).
cnf(c_0_10, plain, (r2_hidden(k4_tarski(X1,esk1_5(X2,X3,X4,X1,X5)),X2)|~r2_hidden(k4_tarski(X1,X5),X4)|X4!=k5_relat_1(X2,X3)|~v1_relat_1(X4)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (v1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_12, plain, ![X25, X26, X27, X29, X30, X31, X32, X34]:(((~r2_hidden(X27,X26)|r2_hidden(k4_tarski(X27,esk5_3(X25,X26,X27)),X25)|X26!=k1_relat_1(X25))&(~r2_hidden(k4_tarski(X29,X30),X25)|r2_hidden(X29,X26)|X26!=k1_relat_1(X25)))&((~r2_hidden(esk6_2(X31,X32),X32)|~r2_hidden(k4_tarski(esk6_2(X31,X32),X34),X31)|X32=k1_relat_1(X31))&(r2_hidden(esk6_2(X31,X32),X32)|r2_hidden(k4_tarski(esk6_2(X31,X32),esk7_2(X31,X32)),X31)|X32=k1_relat_1(X31)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
cnf(c_0_13, plain, (r2_hidden(k4_tarski(esk1_5(X1,X2,X3,X4,X5),X5),X2)|~r2_hidden(k4_tarski(X4,X5),X3)|X3!=k5_relat_1(X1,X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14, plain, (X1=k1_funct_1(X2,X3)|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (r2_hidden(k4_tarski(X1,esk1_5(X2,X3,k5_relat_1(X2,X3),X1,X4)),X2)|~r2_hidden(k4_tarski(X1,X4),k5_relat_1(X2,X3))|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_10]), c_0_11])).
fof(c_0_16, negated_conjecture, ~(![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(k1_relat_1(k5_relat_1(X2,X1))=k1_relat_1(X2)=>r1_tarski(k2_relat_1(X2),k1_relat_1(X1)))))), inference(assume_negation,[status(cth)],[t27_funct_1])).
cnf(c_0_17, plain, (r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, plain, (r2_hidden(k4_tarski(esk1_5(X1,X2,k5_relat_1(X1,X2),X3,X4),X4),X2)|~r2_hidden(k4_tarski(X3,X4),k5_relat_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_13]), c_0_11])).
cnf(c_0_19, plain, (esk1_5(X1,X2,k5_relat_1(X1,X2),X3,X4)=k1_funct_1(X1,X3)|~v1_funct_1(X1)|~r2_hidden(k4_tarski(X3,X4),k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_20, plain, (r2_hidden(k4_tarski(X1,esk5_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_21, negated_conjecture, ((v1_relat_1(esk8_0)&v1_funct_1(esk8_0))&((v1_relat_1(esk9_0)&v1_funct_1(esk9_0))&(k1_relat_1(k5_relat_1(esk9_0,esk8_0))=k1_relat_1(esk9_0)&~r1_tarski(k2_relat_1(esk9_0),k1_relat_1(esk8_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).
fof(c_0_22, plain, ![X38, X39, X40, X42, X43, X44, X45, X47]:(((~r2_hidden(X40,X39)|r2_hidden(k4_tarski(esk10_3(X38,X39,X40),X40),X38)|X39!=k2_relat_1(X38))&(~r2_hidden(k4_tarski(X43,X42),X38)|r2_hidden(X42,X39)|X39!=k2_relat_1(X38)))&((~r2_hidden(esk11_2(X44,X45),X45)|~r2_hidden(k4_tarski(X47,esk11_2(X44,X45)),X44)|X45=k2_relat_1(X44))&(r2_hidden(esk11_2(X44,X45),X45)|r2_hidden(k4_tarski(esk12_2(X44,X45),esk11_2(X44,X45)),X44)|X45=k2_relat_1(X44)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_23, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (r2_hidden(k4_tarski(k1_funct_1(X1,X2),X3),X4)|~v1_funct_1(X1)|~r2_hidden(k4_tarski(X2,X3),k5_relat_1(X1,X4))|~v1_relat_1(X4)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25, plain, (r2_hidden(k4_tarski(X1,esk5_3(X2,k1_relat_1(X2),X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_26, negated_conjecture, (k1_relat_1(k5_relat_1(esk9_0,esk8_0))=k1_relat_1(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (r2_hidden(k4_tarski(esk10_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, plain, (r2_hidden(k1_funct_1(X1,X2),k1_relat_1(X3))|~v1_funct_1(X1)|~r2_hidden(k4_tarski(X2,X4),k5_relat_1(X1,X3))|~v1_relat_1(X3)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_29, negated_conjecture, (r2_hidden(k4_tarski(X1,esk5_3(k5_relat_1(esk9_0,esk8_0),k1_relat_1(esk9_0),X1)),k5_relat_1(esk9_0,esk8_0))|~r2_hidden(X1,k1_relat_1(esk9_0))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_30, negated_conjecture, (v1_funct_1(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, negated_conjecture, (v1_relat_1(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_32, negated_conjecture, (v1_relat_1(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_33, plain, (r2_hidden(k4_tarski(esk10_3(X1,k2_relat_1(X1),X2),X2),X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_34, negated_conjecture, (r2_hidden(k1_funct_1(esk9_0,X1),k1_relat_1(esk8_0))|~r2_hidden(X1,k1_relat_1(esk9_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30]), c_0_31]), c_0_32])])).
cnf(c_0_35, plain, (k1_funct_1(X1,esk10_3(X1,k2_relat_1(X1),X2))=X2|~v1_funct_1(X1)|~r2_hidden(X2,k2_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_14, c_0_33])).
cnf(c_0_36, plain, (r2_hidden(esk10_3(X1,k2_relat_1(X1),X2),k1_relat_1(X1))|~r2_hidden(X2,k2_relat_1(X1))), inference(spm,[status(thm)],[c_0_23, c_0_33])).
fof(c_0_37, plain, ![X49, X50, X51, X52, X53]:((~r1_tarski(X49,X50)|(~r2_hidden(X51,X49)|r2_hidden(X51,X50)))&((r2_hidden(esk13_2(X52,X53),X52)|r1_tarski(X52,X53))&(~r2_hidden(esk13_2(X52,X53),X53)|r1_tarski(X52,X53)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_38, negated_conjecture, (r2_hidden(X1,k1_relat_1(esk8_0))|~r2_hidden(X1,k2_relat_1(esk9_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_30]), c_0_32])]), c_0_36])).
cnf(c_0_39, plain, (r2_hidden(esk13_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_40, plain, (r1_tarski(X1,X2)|~r2_hidden(esk13_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_41, negated_conjecture, (r1_tarski(k2_relat_1(esk9_0),X1)|r2_hidden(esk13_2(k2_relat_1(esk9_0),X1),k1_relat_1(esk8_0))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_42, negated_conjecture, (~r1_tarski(k2_relat_1(esk9_0),k1_relat_1(esk8_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_43, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 44
# Proof object clause steps            : 29
# Proof object formula steps           : 15
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 16
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 189
# ...of these trivial                  : 0
# ...subsumed                          : 41
# ...remaining for further processing  : 148
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 552
# ...of the previous two non-trivial   : 530
# Contextual simplify-reflections      : 10
# Paramodulations                      : 540
# Factorizations                       : 4
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 113
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 106
# Current number of unprocessed clauses: 394
# ...number of literals in the above   : 2445
# Current number of archived formulas  : 0
# Current number of archived clauses   : 27
# Clause-clause subsumption calls (NU) : 3825
# Rec. Clause-clause subsumption calls : 627
# Non-unit clause-clause subsumptions  : 52
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 17220

# -------------------------------------------------
# User time                : 0.017 s
# System time              : 0.003 s
# Total time               : 0.020 s
# Maximum resident set size: 28744 pages
