# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t64_zfmisc_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k4_xboole_0(X2,k1_tarski(X3)))<=>(r2_hidden(X1,X2)&~(X1=X3))), file('data/MPTP/proofs/18f3d1d6f72a4e34b227d3f4dcdee234.E_input', t64_zfmisc_1)).
fof(t76_enumset1, axiom, ![X1]:k1_enumset1(X1,X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/18f3d1d6f72a4e34b227d3f4dcdee234.E_input', t76_enumset1)).
fof(t89_enumset1, axiom, ![X1, X2, X3]:k5_enumset1(X1,X1,X1,X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/18f3d1d6f72a4e34b227d3f4dcdee234.E_input', t89_enumset1)).
fof(t46_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0, file('data/MPTP/proofs/18f3d1d6f72a4e34b227d3f4dcdee234.E_input', t46_xboole_1)).
fof(t22_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/18f3d1d6f72a4e34b227d3f4dcdee234.E_input', t22_xboole_1)).
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/18f3d1d6f72a4e34b227d3f4dcdee234.E_input', d2_zfmisc_1)).
fof(t50_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3)), file('data/MPTP/proofs/18f3d1d6f72a4e34b227d3f4dcdee234.E_input', t50_xboole_1)).
fof(t110_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>k7_relat_1(X1,k1_xboole_0)=k1_xboole_0), file('data/MPTP/proofs/18f3d1d6f72a4e34b227d3f4dcdee234.E_input', t110_relat_1)).
fof(t96_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k7_relat_1(X2,X1)=k3_xboole_0(X2,k2_zfmisc_1(X1,k2_relat_1(X2)))), file('data/MPTP/proofs/18f3d1d6f72a4e34b227d3f4dcdee234.E_input', t96_relat_1)).
fof(c_0_9, plain, ![X7, X8, X9]:(((r2_hidden(X7,X8)|~r2_hidden(X7,k4_xboole_0(X8,k1_tarski(X9))))&(X7!=X9|~r2_hidden(X7,k4_xboole_0(X8,k1_tarski(X9)))))&(~r2_hidden(X7,X8)|X7=X9|r2_hidden(X7,k4_xboole_0(X8,k1_tarski(X9))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t64_zfmisc_1])])])])).
fof(c_0_10, plain, ![X10]:k1_enumset1(X10,X10,X10)=k1_tarski(X10), inference(variable_rename,[status(thm)],[t76_enumset1])).
fof(c_0_11, plain, ![X11, X12, X13]:k5_enumset1(X11,X11,X11,X11,X11,X12,X13)=k1_enumset1(X11,X12,X13), inference(variable_rename,[status(thm)],[t89_enumset1])).
cnf(c_0_12, plain, (X1!=X2|~r2_hidden(X1,k4_xboole_0(X3,k1_tarski(X2)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (k1_enumset1(X1,X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (k5_enumset1(X1,X1,X1,X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_15, plain, ![X14, X15]:k4_xboole_0(X14,k2_xboole_0(X14,X15))=k1_xboole_0, inference(variable_rename,[status(thm)],[t46_xboole_1])).
fof(c_0_16, plain, ![X16, X17]:k2_xboole_0(X16,k3_xboole_0(X16,X17))=X16, inference(variable_rename,[status(thm)],[t22_xboole_1])).
cnf(c_0_17, plain, (X1!=X2|~r2_hidden(X1,k4_xboole_0(X3,k5_enumset1(X2,X2,X2,X2,X2,X2,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_13]), c_0_14])).
cnf(c_0_18, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, (~r2_hidden(X1,k4_xboole_0(X2,k5_enumset1(X1,X1,X1,X1,X1,X1,X1)))), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_21, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_18, c_0_19])).
fof(c_0_22, plain, ![X18, X19, X20, X21, X24, X25, X26, X27, X28, X29, X31, X32]:((((r2_hidden(esk1_4(X18,X19,X20,X21),X18)|~r2_hidden(X21,X20)|X20!=k2_zfmisc_1(X18,X19))&((r2_hidden(esk2_4(X18,X19,X20,X21),X19)|~r2_hidden(X21,X20)|X20!=k2_zfmisc_1(X18,X19))&(X21=k4_tarski(esk1_4(X18,X19,X20,X21),esk2_4(X18,X19,X20,X21))|~r2_hidden(X21,X20)|X20!=k2_zfmisc_1(X18,X19))))&(~r2_hidden(X25,X18)|(~r2_hidden(X26,X19)|X24!=k4_tarski(X25,X26))|r2_hidden(X24,X20)|X20!=k2_zfmisc_1(X18,X19)))&((~r2_hidden(esk3_3(X27,X28,X29),X29)|(~r2_hidden(X31,X27)|(~r2_hidden(X32,X28)|esk3_3(X27,X28,X29)!=k4_tarski(X31,X32)))|X29=k2_zfmisc_1(X27,X28))&((r2_hidden(esk4_3(X27,X28,X29),X27)|r2_hidden(esk3_3(X27,X28,X29),X29)|X29=k2_zfmisc_1(X27,X28))&((r2_hidden(esk5_3(X27,X28,X29),X28)|r2_hidden(esk3_3(X27,X28,X29),X29)|X29=k2_zfmisc_1(X27,X28))&(esk3_3(X27,X28,X29)=k4_tarski(esk4_3(X27,X28,X29),esk5_3(X27,X28,X29))|r2_hidden(esk3_3(X27,X28,X29),X29)|X29=k2_zfmisc_1(X27,X28)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
cnf(c_0_23, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_24, plain, (r2_hidden(esk4_3(X1,X2,X3),X1)|r2_hidden(esk3_3(X1,X2,X3),X3)|X3=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_25, plain, ![X35, X36, X37]:k3_xboole_0(X35,k4_xboole_0(X36,X37))=k4_xboole_0(k3_xboole_0(X35,X36),k3_xboole_0(X35,X37)), inference(variable_rename,[status(thm)],[t50_xboole_1])).
fof(c_0_26, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>k7_relat_1(X1,k1_xboole_0)=k1_xboole_0)), inference(assume_negation,[status(cth)],[t110_relat_1])).
fof(c_0_27, plain, ![X39, X40]:(~v1_relat_1(X40)|k7_relat_1(X40,X39)=k3_xboole_0(X40,k2_zfmisc_1(X39,k2_relat_1(X40)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t96_relat_1])])).
cnf(c_0_28, plain, (X1=k2_zfmisc_1(k1_xboole_0,X2)|r2_hidden(esk3_3(k1_xboole_0,X2,X1),X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_29, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_30, negated_conjecture, (v1_relat_1(esk6_0)&k7_relat_1(esk6_0,k1_xboole_0)!=k1_xboole_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])).
cnf(c_0_31, plain, (k7_relat_1(X1,X2)=k3_xboole_0(X1,k2_zfmisc_1(X2,k2_relat_1(X1)))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_32, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_23, c_0_28])).
cnf(c_0_33, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_29]), c_0_21])).
cnf(c_0_34, negated_conjecture, (k7_relat_1(esk6_0,k1_xboole_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_35, plain, (k7_relat_1(X1,k1_xboole_0)=k1_xboole_0|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])).
cnf(c_0_36, negated_conjecture, (v1_relat_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 38
# Proof object clause steps            : 19
# Proof object formula steps           : 19
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 7
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 17
# Processed clauses                    : 55
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 55
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 108
# ...of the previous two non-trivial   : 54
# Contextual simplify-reflections      : 0
# Paramodulations                      : 103
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 32
#    Positive orientable unit clauses  : 18
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 11
# Current number of unprocessed clauses: 33
# ...number of literals in the above   : 65
# Current number of archived formulas  : 0
# Current number of archived clauses   : 20
# Clause-clause subsumption calls (NU) : 33
# Rec. Clause-clause subsumption calls : 16
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 5
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2136

# -------------------------------------------------
# User time                : 0.003 s
# System time              : 0.006 s
# Total time               : 0.010 s
# Maximum resident set size: 28716 pages
