# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t45_wellord1, conjecture, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r3_wellord1(X1,X2,X3)=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X1)=>(X4=X5|(r2_hidden(k4_tarski(k1_funct_1(X3,X4),k1_funct_1(X3,X5)),X2)&~(k1_funct_1(X3,X4)=k1_funct_1(X3,X5))))))))), file('data/MPTP/proofs/19468c1fd16244f1b62d4d915bcf5904.E_input', t45_wellord1)).
fof(d7_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r3_wellord1(X1,X2,X3)<=>(k1_relat_1(X3)=k3_relat_1(X1)&(k2_relat_1(X3)=k3_relat_1(X2)&(v2_funct_1(X3)&![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X1)<=>(r2_hidden(X4,k3_relat_1(X1))&(r2_hidden(X5,k3_relat_1(X1))&r2_hidden(k4_tarski(k1_funct_1(X3,X4),k1_funct_1(X3,X5)),X2))))))))))), file('data/MPTP/proofs/19468c1fd16244f1b62d4d915bcf5904.E_input', d7_wellord1)).
fof(d8_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)<=>![X2, X3]:((r2_hidden(X2,k1_relat_1(X1))&(r2_hidden(X3,k1_relat_1(X1))&k1_funct_1(X1,X2)=k1_funct_1(X1,X3)))=>X2=X3))), file('data/MPTP/proofs/19468c1fd16244f1b62d4d915bcf5904.E_input', d8_funct_1)).
fof(c_0_3, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r3_wellord1(X1,X2,X3)=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X1)=>(X4=X5|(r2_hidden(k4_tarski(k1_funct_1(X3,X4),k1_funct_1(X3,X5)),X2)&~(k1_funct_1(X3,X4)=k1_funct_1(X3,X5)))))))))), inference(assume_negation,[status(cth)],[t45_wellord1])).
fof(c_0_4, plain, ![X16, X17, X18, X19, X20, X21, X22]:(((k1_relat_1(X18)=k3_relat_1(X16)|~r3_wellord1(X16,X17,X18)|(~v1_relat_1(X18)|~v1_funct_1(X18))|~v1_relat_1(X17)|~v1_relat_1(X16))&((k2_relat_1(X18)=k3_relat_1(X17)|~r3_wellord1(X16,X17,X18)|(~v1_relat_1(X18)|~v1_funct_1(X18))|~v1_relat_1(X17)|~v1_relat_1(X16))&((v2_funct_1(X18)|~r3_wellord1(X16,X17,X18)|(~v1_relat_1(X18)|~v1_funct_1(X18))|~v1_relat_1(X17)|~v1_relat_1(X16))&(((r2_hidden(X19,k3_relat_1(X16))|~r2_hidden(k4_tarski(X19,X20),X16)|~r3_wellord1(X16,X17,X18)|(~v1_relat_1(X18)|~v1_funct_1(X18))|~v1_relat_1(X17)|~v1_relat_1(X16))&((r2_hidden(X20,k3_relat_1(X16))|~r2_hidden(k4_tarski(X19,X20),X16)|~r3_wellord1(X16,X17,X18)|(~v1_relat_1(X18)|~v1_funct_1(X18))|~v1_relat_1(X17)|~v1_relat_1(X16))&(r2_hidden(k4_tarski(k1_funct_1(X18,X19),k1_funct_1(X18,X20)),X17)|~r2_hidden(k4_tarski(X19,X20),X16)|~r3_wellord1(X16,X17,X18)|(~v1_relat_1(X18)|~v1_funct_1(X18))|~v1_relat_1(X17)|~v1_relat_1(X16))))&(~r2_hidden(X21,k3_relat_1(X16))|(~r2_hidden(X22,k3_relat_1(X16))|~r2_hidden(k4_tarski(k1_funct_1(X18,X21),k1_funct_1(X18,X22)),X17))|r2_hidden(k4_tarski(X21,X22),X16)|~r3_wellord1(X16,X17,X18)|(~v1_relat_1(X18)|~v1_funct_1(X18))|~v1_relat_1(X17)|~v1_relat_1(X16))))))&((~r2_hidden(k4_tarski(esk8_3(X16,X17,X18),esk9_3(X16,X17,X18)),X16)|(~r2_hidden(esk8_3(X16,X17,X18),k3_relat_1(X16))|(~r2_hidden(esk9_3(X16,X17,X18),k3_relat_1(X16))|~r2_hidden(k4_tarski(k1_funct_1(X18,esk8_3(X16,X17,X18)),k1_funct_1(X18,esk9_3(X16,X17,X18))),X17)))|~v2_funct_1(X18)|k2_relat_1(X18)!=k3_relat_1(X17)|k1_relat_1(X18)!=k3_relat_1(X16)|r3_wellord1(X16,X17,X18)|(~v1_relat_1(X18)|~v1_funct_1(X18))|~v1_relat_1(X17)|~v1_relat_1(X16))&((r2_hidden(esk8_3(X16,X17,X18),k3_relat_1(X16))|r2_hidden(k4_tarski(esk8_3(X16,X17,X18),esk9_3(X16,X17,X18)),X16)|~v2_funct_1(X18)|k2_relat_1(X18)!=k3_relat_1(X17)|k1_relat_1(X18)!=k3_relat_1(X16)|r3_wellord1(X16,X17,X18)|(~v1_relat_1(X18)|~v1_funct_1(X18))|~v1_relat_1(X17)|~v1_relat_1(X16))&((r2_hidden(esk9_3(X16,X17,X18),k3_relat_1(X16))|r2_hidden(k4_tarski(esk8_3(X16,X17,X18),esk9_3(X16,X17,X18)),X16)|~v2_funct_1(X18)|k2_relat_1(X18)!=k3_relat_1(X17)|k1_relat_1(X18)!=k3_relat_1(X16)|r3_wellord1(X16,X17,X18)|(~v1_relat_1(X18)|~v1_funct_1(X18))|~v1_relat_1(X17)|~v1_relat_1(X16))&(r2_hidden(k4_tarski(k1_funct_1(X18,esk8_3(X16,X17,X18)),k1_funct_1(X18,esk9_3(X16,X17,X18))),X17)|r2_hidden(k4_tarski(esk8_3(X16,X17,X18),esk9_3(X16,X17,X18)),X16)|~v2_funct_1(X18)|k2_relat_1(X18)!=k3_relat_1(X17)|k1_relat_1(X18)!=k3_relat_1(X16)|r3_wellord1(X16,X17,X18)|(~v1_relat_1(X18)|~v1_funct_1(X18))|~v1_relat_1(X17)|~v1_relat_1(X16)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_wellord1])])])])])])).
fof(c_0_5, negated_conjecture, (v1_relat_1(esk1_0)&(v1_relat_1(esk2_0)&((v1_relat_1(esk3_0)&v1_funct_1(esk3_0))&(r3_wellord1(esk1_0,esk2_0,esk3_0)&(r2_hidden(k4_tarski(esk4_0,esk5_0),esk1_0)&(esk4_0!=esk5_0&(~r2_hidden(k4_tarski(k1_funct_1(esk3_0,esk4_0),k1_funct_1(esk3_0,esk5_0)),esk2_0)|k1_funct_1(esk3_0,esk4_0)=k1_funct_1(esk3_0,esk5_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_3])])])])).
cnf(c_0_6, plain, (r2_hidden(k4_tarski(k1_funct_1(X1,X2),k1_funct_1(X1,X3)),X4)|~r2_hidden(k4_tarski(X2,X3),X5)|~r3_wellord1(X5,X4,X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X4)|~v1_relat_1(X5)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, negated_conjecture, (r2_hidden(k4_tarski(esk4_0,esk5_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (r2_hidden(X1,k3_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)|~r3_wellord1(X2,X4,X5)|~v1_relat_1(X5)|~v1_funct_1(X5)|~v1_relat_1(X4)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_10, plain, ![X11, X12, X13]:((~v2_funct_1(X11)|(~r2_hidden(X12,k1_relat_1(X11))|(~r2_hidden(X13,k1_relat_1(X11))|k1_funct_1(X11,X12)!=k1_funct_1(X11,X13))|X12=X13)|(~v1_relat_1(X11)|~v1_funct_1(X11)))&(((r2_hidden(esk6_1(X11),k1_relat_1(X11))|v2_funct_1(X11)|(~v1_relat_1(X11)|~v1_funct_1(X11)))&((r2_hidden(esk7_1(X11),k1_relat_1(X11))|v2_funct_1(X11)|(~v1_relat_1(X11)|~v1_funct_1(X11)))&(k1_funct_1(X11,esk6_1(X11))=k1_funct_1(X11,esk7_1(X11))|v2_funct_1(X11)|(~v1_relat_1(X11)|~v1_funct_1(X11)))))&(esk6_1(X11)!=esk7_1(X11)|v2_funct_1(X11)|(~v1_relat_1(X11)|~v1_funct_1(X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_funct_1])])])])])).
cnf(c_0_11, negated_conjecture, (k1_funct_1(esk3_0,esk4_0)=k1_funct_1(esk3_0,esk5_0)|~r2_hidden(k4_tarski(k1_funct_1(esk3_0,esk4_0),k1_funct_1(esk3_0,esk5_0)),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_12, negated_conjecture, (r2_hidden(k4_tarski(k1_funct_1(X1,esk4_0),k1_funct_1(X1,esk5_0)),X2)|~r3_wellord1(esk1_0,X2,X1)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8])])).
cnf(c_0_13, negated_conjecture, (r3_wellord1(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_14, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_15, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_16, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_17, plain, (v2_funct_1(X1)|~r3_wellord1(X2,X3,X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_18, plain, (k1_relat_1(X1)=k3_relat_1(X2)|~r3_wellord1(X2,X3,X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_19, negated_conjecture, (r2_hidden(esk5_0,k3_relat_1(esk1_0))|~r3_wellord1(esk1_0,X1,X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_7]), c_0_8])])).
cnf(c_0_20, plain, (r2_hidden(X1,k3_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)|~r3_wellord1(X2,X4,X5)|~v1_relat_1(X5)|~v1_funct_1(X5)|~v1_relat_1(X4)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_21, plain, (X2=X3|~v2_funct_1(X1)|~r2_hidden(X2,k1_relat_1(X1))|~r2_hidden(X3,k1_relat_1(X1))|k1_funct_1(X1,X2)!=k1_funct_1(X1,X3)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_22, negated_conjecture, (k1_funct_1(esk3_0,esk5_0)=k1_funct_1(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13]), c_0_14]), c_0_15]), c_0_16])])).
cnf(c_0_23, negated_conjecture, (v2_funct_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_13]), c_0_14]), c_0_15]), c_0_8]), c_0_16])])).
cnf(c_0_24, negated_conjecture, (k1_relat_1(esk3_0)=k3_relat_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_13]), c_0_14]), c_0_15]), c_0_8]), c_0_16])])).
cnf(c_0_25, negated_conjecture, (r2_hidden(esk5_0,k3_relat_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_13]), c_0_14]), c_0_16]), c_0_15])])).
cnf(c_0_26, negated_conjecture, (r2_hidden(esk4_0,k3_relat_1(esk1_0))|~r3_wellord1(esk1_0,X1,X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_7]), c_0_8])])).
cnf(c_0_27, negated_conjecture, (esk5_0=X1|k1_funct_1(esk3_0,esk4_0)!=k1_funct_1(esk3_0,X1)|~r2_hidden(X1,k3_relat_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24]), c_0_24]), c_0_25]), c_0_14]), c_0_16])])).
cnf(c_0_28, negated_conjecture, (r2_hidden(esk4_0,k3_relat_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_13]), c_0_14]), c_0_16]), c_0_15])])).
cnf(c_0_29, negated_conjecture, (esk4_0!=esk5_0), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_30, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_27]), c_0_28])]), c_0_29]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 31
# Proof object clause steps            : 24
# Proof object formula steps           : 7
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 39
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 24
# Processed clauses                    : 58
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 58
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 29
# ...of the previous two non-trivial   : 22
# Contextual simplify-reflections      : 0
# Paramodulations                      : 27
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 31
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 18
# Current number of unprocessed clauses: 12
# ...number of literals in the above   : 89
# Current number of archived formulas  : 0
# Current number of archived clauses   : 27
# Clause-clause subsumption calls (NU) : 670
# Rec. Clause-clause subsumption calls : 16
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 13
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3545

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28716 pages
