# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t36_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>r1_lattice3(X1,k4_waybel_0(X1,X2),X3))))), file('data/MPTP/proofs/1a7ac80ae2174f06a1cccf278607bc4d.E_input', t36_waybel_0)).
fof(t15_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k4_waybel_0(X1,X2)=a_2_1_waybel_0(X1,X2))), file('data/MPTP/proofs/1a7ac80ae2174f06a1cccf278607bc4d.E_input', t15_waybel_0)).
fof(fraenkel_a_2_1_waybel_0, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&l1_orders_2(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r2_hidden(X1,a_2_1_waybel_0(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X5,X4)&r2_hidden(X5,X3))))))), file('data/MPTP/proofs/1a7ac80ae2174f06a1cccf278607bc4d.E_input', fraenkel_a_2_1_waybel_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/1a7ac80ae2174f06a1cccf278607bc4d.E_input', t24_orders_2)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/1a7ac80ae2174f06a1cccf278607bc4d.E_input', d8_lattice3)).
fof(t26_orders_2, axiom, ![X1]:((v4_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X3,X4))=>r1_orders_2(X1,X2,X4)))))), file('data/MPTP/proofs/1a7ac80ae2174f06a1cccf278607bc4d.E_input', t26_orders_2)).
fof(c_0_6, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>r1_lattice3(X1,k4_waybel_0(X1,X2),X3)))))), inference(assume_negation,[status(cth)],[t36_waybel_0])).
fof(c_0_7, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&((~r1_lattice3(esk1_0,esk2_0,esk3_0)|~r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0))&(r1_lattice3(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_6])])])])).
fof(c_0_8, plain, ![X21, X22]:(v2_struct_0(X21)|~l1_orders_2(X21)|(~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))|k4_waybel_0(X21,X22)=a_2_1_waybel_0(X21,X22))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t15_waybel_0])])])])).
fof(c_0_9, plain, ![X14, X15, X16, X19, X20]:(((m1_subset_1(esk5_3(X14,X15,X16),u1_struct_0(X15))|~r2_hidden(X14,a_2_1_waybel_0(X15,X16))|(v2_struct_0(X15)|~l1_orders_2(X15)|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X15)))))&((X14=esk5_3(X14,X15,X16)|~r2_hidden(X14,a_2_1_waybel_0(X15,X16))|(v2_struct_0(X15)|~l1_orders_2(X15)|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X15)))))&((m1_subset_1(esk6_3(X14,X15,X16),u1_struct_0(X15))|~r2_hidden(X14,a_2_1_waybel_0(X15,X16))|(v2_struct_0(X15)|~l1_orders_2(X15)|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X15)))))&((r1_orders_2(X15,esk6_3(X14,X15,X16),esk5_3(X14,X15,X16))|~r2_hidden(X14,a_2_1_waybel_0(X15,X16))|(v2_struct_0(X15)|~l1_orders_2(X15)|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X15)))))&(r2_hidden(esk6_3(X14,X15,X16),X16)|~r2_hidden(X14,a_2_1_waybel_0(X15,X16))|(v2_struct_0(X15)|~l1_orders_2(X15)|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X15)))))))))&(~m1_subset_1(X19,u1_struct_0(X15))|(X14!=X19|(~m1_subset_1(X20,u1_struct_0(X15))|(~r1_orders_2(X15,X20,X19)|~r2_hidden(X20,X16))))|r2_hidden(X14,a_2_1_waybel_0(X15,X16))|(v2_struct_0(X15)|~l1_orders_2(X15)|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X15)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_1_waybel_0])])])])])])).
fof(c_0_10, plain, ![X27, X28]:(v2_struct_0(X27)|(~v3_orders_2(X27)|~l1_orders_2(X27))|(~m1_subset_1(X28,u1_struct_0(X27))|r1_orders_2(X27,X28,X28))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_11, plain, ![X9, X10, X11, X12]:((~r1_lattice3(X9,X10,X11)|(~m1_subset_1(X12,u1_struct_0(X9))|(~r2_hidden(X12,X10)|r1_orders_2(X9,X11,X12)))|~m1_subset_1(X11,u1_struct_0(X9))|~l1_orders_2(X9))&((m1_subset_1(esk4_3(X9,X10,X11),u1_struct_0(X9))|r1_lattice3(X9,X10,X11)|~m1_subset_1(X11,u1_struct_0(X9))|~l1_orders_2(X9))&((r2_hidden(esk4_3(X9,X10,X11),X10)|r1_lattice3(X9,X10,X11)|~m1_subset_1(X11,u1_struct_0(X9))|~l1_orders_2(X9))&(~r1_orders_2(X9,X11,esk4_3(X9,X10,X11))|r1_lattice3(X9,X10,X11)|~m1_subset_1(X11,u1_struct_0(X9))|~l1_orders_2(X9))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
cnf(c_0_12, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, plain, (v2_struct_0(X1)|k4_waybel_0(X1,X2)=a_2_1_waybel_0(X1,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_15, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_16, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_17, plain, (r2_hidden(X3,a_2_1_waybel_0(X2,X5))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~m1_subset_1(X4,u1_struct_0(X2))|~r1_orders_2(X2,X4,X1)|~r2_hidden(X4,X5)|~l1_orders_2(X2)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_18, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_20, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_21, negated_conjecture, (r1_lattice3(esk1_0,a_2_1_waybel_0(esk1_0,esk2_0),esk3_0)|r1_lattice3(esk1_0,esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_22, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_23, plain, (r2_hidden(X1,a_2_1_waybel_0(X2,X3))|v2_struct_0(X2)|~r1_orders_2(X2,X4,X1)|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_24, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_15])]), c_0_16])).
cnf(c_0_25, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,X1)|r1_lattice3(esk1_0,esk2_0,esk3_0)|~r2_hidden(X1,a_2_1_waybel_0(esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_15])])).
cnf(c_0_26, negated_conjecture, (r2_hidden(X1,a_2_1_waybel_0(esk1_0,X2))|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_15])]), c_0_16])).
cnf(c_0_27, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,X1)|r1_lattice3(esk1_0,esk2_0,esk3_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_14])])).
cnf(c_0_28, plain, (r2_hidden(esk4_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_29, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk4_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_30, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk4_3(X1,esk2_0,X2))|r1_lattice3(esk1_0,esk2_0,esk3_0)|r1_lattice3(X1,esk2_0,X2)|~m1_subset_1(esk4_3(X1,esk2_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_31, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk3_0)|~m1_subset_1(esk4_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_22]), c_0_15])])).
cnf(c_0_32, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_33, plain, ![X23, X24, X25, X26]:(~v4_orders_2(X23)|~l1_orders_2(X23)|(~m1_subset_1(X24,u1_struct_0(X23))|(~m1_subset_1(X25,u1_struct_0(X23))|(~m1_subset_1(X26,u1_struct_0(X23))|(~r1_orders_2(X23,X24,X25)|~r1_orders_2(X23,X25,X26)|r1_orders_2(X23,X24,X26)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t26_orders_2])])])).
cnf(c_0_34, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_22]), c_0_15])])).
cnf(c_0_35, plain, (r1_orders_2(X1,X2,X4)|~v4_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_36, plain, (r1_orders_2(X1,esk6_3(X2,X1,X3),esk5_3(X2,X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,a_2_1_waybel_0(X1,X3))|~l1_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_37, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_38, plain, (m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_39, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,X1)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_34]), c_0_22]), c_0_15])])).
cnf(c_0_40, plain, (r2_hidden(esk6_3(X1,X2,X3),X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_41, plain, (r1_orders_2(X1,X2,esk5_3(X3,X1,X4))|v2_struct_0(X1)|~r1_orders_2(X1,X2,esk6_3(X3,X1,X4))|~r2_hidden(X3,a_2_1_waybel_0(X1,X4))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v4_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37]), c_0_38])).
cnf(c_0_42, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk6_3(X1,X2,esk2_0))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,esk2_0))|~m1_subset_1(esk6_3(X1,X2,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_43, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_44, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk5_3(X1,esk1_0,esk2_0))|~r2_hidden(X1,a_2_1_waybel_0(esk1_0,esk2_0))|~m1_subset_1(esk6_3(X1,esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_14]), c_0_22]), c_0_15]), c_0_43])]), c_0_16])).
cnf(c_0_45, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk5_3(X1,esk1_0,esk2_0))|~r2_hidden(X1,a_2_1_waybel_0(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_37]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_46, plain, (X1=esk5_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_47, negated_conjecture, (~r1_lattice3(esk1_0,esk2_0,esk3_0)|~r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_48, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,X1)|~r2_hidden(X1,a_2_1_waybel_0(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_49, negated_conjecture, (~r1_lattice3(esk1_0,a_2_1_waybel_0(esk1_0,esk2_0),esk3_0)|~r1_lattice3(esk1_0,esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_13]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_50, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk4_3(X1,a_2_1_waybel_0(esk1_0,esk2_0),X2))|r1_lattice3(X1,a_2_1_waybel_0(esk1_0,esk2_0),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_48, c_0_28])).
cnf(c_0_51, negated_conjecture, (~r1_lattice3(esk1_0,a_2_1_waybel_0(esk1_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_34])])).
cnf(c_0_52, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_50]), c_0_22]), c_0_15])]), c_0_51]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 53
# Proof object clause steps            : 40
# Proof object formula steps           : 13
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 51
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 21
# Processed clauses                    : 106
# ...of these trivial                  : 0
# ...subsumed                          : 14
# ...remaining for further processing  : 92
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 7
# Backward-rewritten                   : 10
# Generated clauses                    : 110
# ...of the previous two non-trivial   : 102
# Contextual simplify-reflections      : 4
# Paramodulations                      : 109
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 53
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 44
# Current number of unprocessed clauses: 28
# ...number of literals in the above   : 225
# Current number of archived formulas  : 0
# Current number of archived clauses   : 38
# Clause-clause subsumption calls (NU) : 1228
# Rec. Clause-clause subsumption calls : 169
# Non-unit clause-clause subsumptions  : 24
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6748

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.003 s
# Total time               : 0.013 s
# Maximum resident set size: 28656 pages
