# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t53_relset_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(~(v1_xboole_0(X3))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))=>![X5]:(m1_subset_1(X5,X1)=>(r2_hidden(X5,k8_relset_1(X1,X3,X4,X2))<=>?[X6]:(m1_subset_1(X6,X3)&(r2_hidden(k4_tarski(X5,X6),X4)&r2_hidden(X6,X2))))))))), file('data/MPTP/proofs/1add995e7d2e417cb15a2cd902762a7b.E_input', t53_relset_1)).
fof(redefinition_k8_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k8_relset_1(X1,X2,X3,X4)=k10_relat_1(X3,X4)), file('data/MPTP/proofs/1add995e7d2e417cb15a2cd902762a7b.E_input', redefinition_k8_relset_1)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/1add995e7d2e417cb15a2cd902762a7b.E_input', l3_subset_1)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/1add995e7d2e417cb15a2cd902762a7b.E_input', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/1add995e7d2e417cb15a2cd902762a7b.E_input', fc6_relat_1)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/1add995e7d2e417cb15a2cd902762a7b.E_input', l54_zfmisc_1)).
fof(t166_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k10_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k2_relat_1(X3))&(r2_hidden(k4_tarski(X1,X4),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/1add995e7d2e417cb15a2cd902762a7b.E_input', t166_relat_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/1add995e7d2e417cb15a2cd902762a7b.E_input', t1_subset)).
fof(t20_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)=>(r2_hidden(X1,k1_relat_1(X3))&r2_hidden(X2,k2_relat_1(X3))))), file('data/MPTP/proofs/1add995e7d2e417cb15a2cd902762a7b.E_input', t20_relat_1)).
fof(c_0_9, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(~(v1_xboole_0(X3))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))=>![X5]:(m1_subset_1(X5,X1)=>(r2_hidden(X5,k8_relset_1(X1,X3,X4,X2))<=>?[X6]:(m1_subset_1(X6,X3)&(r2_hidden(k4_tarski(X5,X6),X4)&r2_hidden(X6,X2)))))))))), inference(assume_negation,[status(cth)],[t53_relset_1])).
fof(c_0_10, plain, ![X14, X15, X16, X17]:(~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15)))|k8_relset_1(X14,X15,X16,X17)=k10_relat_1(X16,X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_relset_1])])).
fof(c_0_11, negated_conjecture, ![X12]:(~v1_xboole_0(esk1_0)&(~v1_xboole_0(esk2_0)&(~v1_xboole_0(esk3_0)&(m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk3_0)))&(m1_subset_1(esk5_0,esk1_0)&((~r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0))|(~m1_subset_1(X12,esk3_0)|(~r2_hidden(k4_tarski(esk5_0,X12),esk4_0)|~r2_hidden(X12,esk2_0))))&((m1_subset_1(esk6_0,esk3_0)|r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0)))&((r2_hidden(k4_tarski(esk5_0,esk6_0),esk4_0)|r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0)))&(r2_hidden(esk6_0,esk2_0)|r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])])])).
cnf(c_0_12, plain, (k8_relset_1(X2,X3,X1,X4)=k10_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_13, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_14, plain, ![X23, X24, X25]:(~m1_subset_1(X24,k1_zfmisc_1(X23))|(~r2_hidden(X25,X24)|r2_hidden(X25,X23))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
cnf(c_0_15, negated_conjecture, (~r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0))|~m1_subset_1(X1,esk3_0)|~r2_hidden(k4_tarski(esk5_0,X1),esk4_0)|~r2_hidden(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, negated_conjecture, (k8_relset_1(esk1_0,esk3_0,esk4_0,X1)=k10_relat_1(esk4_0,X1)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_17, negated_conjecture, (r2_hidden(k4_tarski(esk5_0,esk6_0),esk4_0)|r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_18, plain, ![X37, X38]:(~v1_relat_1(X37)|(~m1_subset_1(X38,k1_zfmisc_1(X37))|v1_relat_1(X38))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_19, plain, ![X26, X27]:v1_relat_1(k2_zfmisc_1(X26,X27)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_20, plain, ![X31, X32, X33, X34]:(((r2_hidden(X31,X33)|~r2_hidden(k4_tarski(X31,X32),k2_zfmisc_1(X33,X34)))&(r2_hidden(X32,X34)|~r2_hidden(k4_tarski(X31,X32),k2_zfmisc_1(X33,X34))))&(~r2_hidden(X31,X33)|~r2_hidden(X32,X34)|r2_hidden(k4_tarski(X31,X32),k2_zfmisc_1(X33,X34)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
cnf(c_0_21, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, negated_conjecture, (~r2_hidden(esk5_0,k10_relat_1(esk4_0,esk2_0))|~r2_hidden(k4_tarski(esk5_0,X1),esk4_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk3_0)), inference(rw,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_23, negated_conjecture, (r2_hidden(k4_tarski(esk5_0,esk6_0),esk4_0)|r2_hidden(esk5_0,k10_relat_1(esk4_0,esk2_0))), inference(rw,[status(thm)],[c_0_17, c_0_16])).
fof(c_0_24, plain, ![X18, X19, X20, X22]:(((r2_hidden(esk7_3(X18,X19,X20),k2_relat_1(X20))|~r2_hidden(X18,k10_relat_1(X20,X19))|~v1_relat_1(X20))&((r2_hidden(k4_tarski(X18,esk7_3(X18,X19,X20)),X20)|~r2_hidden(X18,k10_relat_1(X20,X19))|~v1_relat_1(X20))&(r2_hidden(esk7_3(X18,X19,X20),X19)|~r2_hidden(X18,k10_relat_1(X20,X19))|~v1_relat_1(X20))))&(~r2_hidden(X22,k2_relat_1(X20))|(~r2_hidden(k4_tarski(X18,X22),X20)|~r2_hidden(X22,X19))|r2_hidden(X18,k10_relat_1(X20,X19))|~v1_relat_1(X20))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t166_relat_1])])])])])).
cnf(c_0_25, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,X2))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(esk1_0,esk3_0))|~r2_hidden(X1,esk4_0)), inference(spm,[status(thm)],[c_0_21, c_0_13])).
cnf(c_0_29, negated_conjecture, (r2_hidden(k4_tarski(esk5_0,esk6_0),esk4_0)|~r2_hidden(k4_tarski(esk5_0,X1),esk4_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk3_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_30, plain, (r2_hidden(k4_tarski(X1,esk7_3(X1,X2,X3)),X3)|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, negated_conjecture, (v1_relat_1(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_13]), c_0_26])])).
fof(c_0_32, plain, ![X35, X36]:(~r2_hidden(X35,X36)|m1_subset_1(X35,X36)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_33, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(k4_tarski(X2,X1),esk4_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34, negated_conjecture, (r2_hidden(esk6_0,esk2_0)|r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_35, plain, ![X28, X29, X30]:((r2_hidden(X28,k1_relat_1(X30))|~r2_hidden(k4_tarski(X28,X29),X30)|~v1_relat_1(X30))&(r2_hidden(X29,k2_relat_1(X30))|~r2_hidden(k4_tarski(X28,X29),X30)|~v1_relat_1(X30))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t20_relat_1])])])).
cnf(c_0_36, negated_conjecture, (r2_hidden(k4_tarski(esk5_0,esk6_0),esk4_0)|~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk2_0)|~r2_hidden(esk5_0,k10_relat_1(esk4_0,X1))|~m1_subset_1(esk7_3(esk5_0,X1,esk4_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])])).
cnf(c_0_37, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, negated_conjecture, (r2_hidden(esk7_3(X1,X2,esk4_0),esk3_0)|~r2_hidden(X1,k10_relat_1(esk4_0,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_30]), c_0_31])])).
cnf(c_0_39, negated_conjecture, (r2_hidden(esk6_0,esk2_0)|~r2_hidden(k4_tarski(esk5_0,X1),esk4_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk3_0)), inference(spm,[status(thm)],[c_0_15, c_0_34])).
cnf(c_0_40, negated_conjecture, (m1_subset_1(esk6_0,esk3_0)|r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_41, plain, (r2_hidden(X3,k10_relat_1(X2,X4))|~r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)|~r2_hidden(X1,X4)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_42, plain, (r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_43, negated_conjecture, (r2_hidden(k4_tarski(esk5_0,esk6_0),esk4_0)|~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk2_0)|~r2_hidden(esk5_0,k10_relat_1(esk4_0,X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_44, plain, (r2_hidden(esk7_3(X1,X2,X3),X2)|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_45, negated_conjecture, (r2_hidden(esk6_0,esk2_0)|~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk2_0)|~r2_hidden(esk5_0,k10_relat_1(esk4_0,X1))|~m1_subset_1(esk7_3(esk5_0,X1,esk4_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_30]), c_0_31])])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(esk6_0,esk3_0)|~r2_hidden(k4_tarski(esk5_0,X1),esk4_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk3_0)), inference(spm,[status(thm)],[c_0_15, c_0_40])).
cnf(c_0_47, plain, (r2_hidden(X1,k10_relat_1(X2,X3))|~v1_relat_1(X2)|~r2_hidden(k4_tarski(X1,X4),X2)|~r2_hidden(X4,X3)), inference(csr,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_48, negated_conjecture, (r2_hidden(k4_tarski(esk5_0,esk6_0),esk4_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_31])]), c_0_23])).
cnf(c_0_49, negated_conjecture, (r2_hidden(esk6_0,esk2_0)|~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk2_0)|~r2_hidden(esk5_0,k10_relat_1(esk4_0,X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_37]), c_0_38])).
cnf(c_0_50, negated_conjecture, (r2_hidden(esk5_0,k10_relat_1(esk4_0,esk2_0))|r2_hidden(esk6_0,esk2_0)), inference(rw,[status(thm)],[c_0_34, c_0_16])).
cnf(c_0_51, negated_conjecture, (m1_subset_1(esk6_0,esk3_0)|~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk2_0)|~r2_hidden(esk5_0,k10_relat_1(esk4_0,X1))|~m1_subset_1(esk7_3(esk5_0,X1,esk4_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_30]), c_0_31])])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk5_0,k10_relat_1(esk4_0,X1))|~r2_hidden(esk6_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_31])])).
cnf(c_0_53, negated_conjecture, (r2_hidden(esk6_0,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_44]), c_0_31])]), c_0_50])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(esk6_0,esk3_0)|~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk2_0)|~r2_hidden(esk5_0,k10_relat_1(esk4_0,X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_37]), c_0_38])).
cnf(c_0_55, negated_conjecture, (r2_hidden(esk5_0,k10_relat_1(esk4_0,esk2_0))|m1_subset_1(esk6_0,esk3_0)), inference(rw,[status(thm)],[c_0_40, c_0_16])).
cnf(c_0_56, negated_conjecture, (~r2_hidden(k4_tarski(esk5_0,X1),esk4_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_52]), c_0_53])])).
cnf(c_0_57, negated_conjecture, (m1_subset_1(esk6_0,esk3_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_44]), c_0_31])]), c_0_55])).
cnf(c_0_58, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_48]), c_0_53]), c_0_57])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 59
# Proof object clause steps            : 40
# Proof object formula steps           : 19
# Proof object conjectures             : 32
# Proof object clause conjectures      : 29
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 34
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 109
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 106
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 19
# Generated clauses                    : 211
# ...of the previous two non-trivial   : 207
# Contextual simplify-reflections      : 9
# Paramodulations                      : 211
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 58
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 42
# Current number of unprocessed clauses: 139
# ...number of literals in the above   : 465
# Current number of archived formulas  : 0
# Current number of archived clauses   : 48
# Clause-clause subsumption calls (NU) : 993
# Rec. Clause-clause subsumption calls : 746
# Non-unit clause-clause subsumptions  : 16
# Unit Clause-clause subsumption calls : 25
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4957

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.006 s
# Total time               : 0.012 s
# Maximum resident set size: 28792 pages
