# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t1_xboole_0, axiom, ![X1, X2, X3]:(r2_hidden(X1,k5_xboole_0(X2,X3))<=>~((r2_hidden(X1,X2)<=>r2_hidden(X1,X3)))), file('data/MPTP/proofs/1ae15767b1354a7d82f6a9b03067c5f3.E_input', t1_xboole_0)).
fof(t5_boole, axiom, ![X1]:k5_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/1ae15767b1354a7d82f6a9b03067c5f3.E_input', t5_boole)).
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/1ae15767b1354a7d82f6a9b03067c5f3.E_input', d2_zfmisc_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/1ae15767b1354a7d82f6a9b03067c5f3.E_input', t7_xboole_0)).
fof(t113_zfmisc_1, conjecture, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/1ae15767b1354a7d82f6a9b03067c5f3.E_input', t113_zfmisc_1)).
fof(c_0_5, plain, ![X10, X11, X12]:(((~r2_hidden(X10,X11)|~r2_hidden(X10,X12)|~r2_hidden(X10,k5_xboole_0(X11,X12)))&(r2_hidden(X10,X11)|r2_hidden(X10,X12)|~r2_hidden(X10,k5_xboole_0(X11,X12))))&((~r2_hidden(X10,X11)|r2_hidden(X10,X12)|r2_hidden(X10,k5_xboole_0(X11,X12)))&(~r2_hidden(X10,X12)|r2_hidden(X10,X11)|r2_hidden(X10,k5_xboole_0(X11,X12))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_0])])])).
fof(c_0_6, plain, ![X9]:k5_xboole_0(X9,k1_xboole_0)=X9, inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_7, plain, (r2_hidden(X1,X3)|r2_hidden(X1,k5_xboole_0(X3,X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r2_hidden(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
fof(c_0_11, plain, ![X15, X16, X17, X18, X21, X22, X23, X24, X25, X26, X28, X29]:((((r2_hidden(esk4_4(X15,X16,X17,X18),X15)|~r2_hidden(X18,X17)|X17!=k2_zfmisc_1(X15,X16))&((r2_hidden(esk5_4(X15,X16,X17,X18),X16)|~r2_hidden(X18,X17)|X17!=k2_zfmisc_1(X15,X16))&(X18=k4_tarski(esk4_4(X15,X16,X17,X18),esk5_4(X15,X16,X17,X18))|~r2_hidden(X18,X17)|X17!=k2_zfmisc_1(X15,X16))))&(~r2_hidden(X22,X15)|(~r2_hidden(X23,X16)|X21!=k4_tarski(X22,X23))|r2_hidden(X21,X17)|X17!=k2_zfmisc_1(X15,X16)))&((~r2_hidden(esk6_3(X24,X25,X26),X26)|(~r2_hidden(X28,X24)|(~r2_hidden(X29,X25)|esk6_3(X24,X25,X26)!=k4_tarski(X28,X29)))|X26=k2_zfmisc_1(X24,X25))&((r2_hidden(esk7_3(X24,X25,X26),X24)|r2_hidden(esk6_3(X24,X25,X26),X26)|X26=k2_zfmisc_1(X24,X25))&((r2_hidden(esk8_3(X24,X25,X26),X25)|r2_hidden(esk6_3(X24,X25,X26),X26)|X26=k2_zfmisc_1(X24,X25))&(esk6_3(X24,X25,X26)=k4_tarski(esk7_3(X24,X25,X26),esk8_3(X24,X25,X26))|r2_hidden(esk6_3(X24,X25,X26),X26)|X26=k2_zfmisc_1(X24,X25)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
cnf(c_0_12, plain, (~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_8]), c_0_10])).
cnf(c_0_13, plain, (r2_hidden(esk8_3(X1,X2,X3),X2)|r2_hidden(esk6_3(X1,X2,X3),X3)|X3=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_14, plain, (r2_hidden(esk5_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (X1=k2_zfmisc_1(X2,k1_xboole_0)|r2_hidden(esk6_3(X2,k1_xboole_0,X1),X1)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_16, plain, (r2_hidden(esk6_3(X1,k1_xboole_0,X2),X2)|~r2_hidden(X3,X2)), inference(er,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_12])])).
cnf(c_0_17, plain, (X1=k2_zfmisc_1(X2,k1_xboole_0)|r2_hidden(esk6_3(X3,k1_xboole_0,X1),X1)), inference(spm,[status(thm)],[c_0_16, c_0_15])).
cnf(c_0_18, plain, (r2_hidden(esk7_3(X1,X2,X3),X1)|r2_hidden(esk6_3(X1,X2,X3),X3)|X3=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_19, plain, ![X13]:(X13=k1_xboole_0|r2_hidden(esk3_1(X13),X13)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_20, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_12, c_0_17])).
cnf(c_0_21, plain, (r2_hidden(esk4_4(X1,X2,X3,X4),X1)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_22, plain, (X1=k2_zfmisc_1(k1_xboole_0,X2)|r2_hidden(esk6_3(k1_xboole_0,X2,X1),X1)), inference(spm,[status(thm)],[c_0_12, c_0_18])).
cnf(c_0_23, plain, (X1=k1_xboole_0|r2_hidden(esk3_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_24, negated_conjecture, ~(![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0))), inference(assume_negation,[status(cth)],[t113_zfmisc_1])).
cnf(c_0_25, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_20]), c_0_12])).
cnf(c_0_26, plain, (r2_hidden(esk4_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X1)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (r2_hidden(esk6_3(k1_xboole_0,X1,X2),X2)|~r2_hidden(X3,X2)), inference(er,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_12])])).
cnf(c_0_28, plain, (X1=k1_xboole_0|r2_hidden(esk6_3(X2,k1_xboole_0,X1),X1)), inference(spm,[status(thm)],[c_0_16, c_0_23])).
fof(c_0_29, negated_conjecture, (((esk1_0!=k1_xboole_0|k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0)&(esk2_0!=k1_xboole_0|k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0))&(k2_zfmisc_1(esk1_0,esk2_0)=k1_xboole_0|(esk1_0=k1_xboole_0|esk2_0=k1_xboole_0))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])])).
cnf(c_0_30, plain, (X1!=k1_xboole_0|~r2_hidden(X2,k2_zfmisc_1(X1,X3))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_31, plain, (X1=k1_xboole_0|r2_hidden(esk6_3(k1_xboole_0,X2,X1),X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_32, plain, (r2_hidden(esk5_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X2)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_14])).
cnf(c_0_33, negated_conjecture, (esk1_0!=k1_xboole_0|k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_35, plain, (X1!=k1_xboole_0|~r2_hidden(X2,k2_zfmisc_1(X3,X1))), inference(spm,[status(thm)],[c_0_25, c_0_32])).
cnf(c_0_36, plain, (r2_hidden(X5,X6)|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)|X5!=k4_tarski(X1,X3)|X6!=k2_zfmisc_1(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_37, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k1_xboole_0|esk1_0=k1_xboole_0|esk2_0=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, negated_conjecture, (esk1_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_39, negated_conjecture, (esk2_0!=k1_xboole_0|k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_40, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X2!=k1_xboole_0), inference(spm,[status(thm)],[c_0_35, c_0_31])).
cnf(c_0_41, plain, (r2_hidden(X1,k2_zfmisc_1(X2,X3))|X1!=k4_tarski(X4,X5)|~r2_hidden(X5,X3)|~r2_hidden(X4,X2)), inference(er,[status(thm)],[c_0_36])).
cnf(c_0_42, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k1_xboole_0|esk2_0=k1_xboole_0), inference(sr,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_43, negated_conjecture, (esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_39, c_0_40])).
fof(c_0_44, plain, (~epred3_0<=>![X2]:~r2_hidden(X2,esk2_0)), introduced(definition)).
fof(c_0_45, plain, (~epred4_0<=>![X1]:~r2_hidden(X1,esk1_0)), introduced(definition)).
cnf(c_0_46, plain, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_47, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k1_xboole_0), inference(sr,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_48, negated_conjecture, (epred3_0|~r2_hidden(X1,esk2_0)), inference(split_equiv,[status(thm)],[c_0_44])).
cnf(c_0_49, negated_conjecture, (~epred4_0|~epred3_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_12]), c_0_44]), c_0_45])).
cnf(c_0_50, negated_conjecture, (epred3_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_31]), c_0_43])).
cnf(c_0_51, negated_conjecture, (~epred4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_50])])).
cnf(c_0_52, negated_conjecture, (~r2_hidden(X1,esk1_0)), inference(sr,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_45]), c_0_51])).
cnf(c_0_53, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_31]), c_0_38]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 54
# Proof object clause steps            : 41
# Proof object formula steps           : 13
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 24
# Proof object simplifying inferences  : 16
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 1165
# ...of these trivial                  : 3
# ...subsumed                          : 885
# ...remaining for further processing  : 277
# Other redundant clauses eliminated   : 216
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 18
# Generated clauses                    : 24517
# ...of the previous two non-trivial   : 21803
# Contextual simplify-reflections      : 5
# Paramodulations                      : 24195
# Factorizations                       : 48
# Equation resolutions                 : 266
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 236
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 223
# Current number of unprocessed clauses: 20504
# ...number of literals in the above   : 128656
# Current number of archived formulas  : 0
# Current number of archived clauses   : 39
# Clause-clause subsumption calls (NU) : 13178
# Rec. Clause-clause subsumption calls : 4854
# Non-unit clause-clause subsumptions  : 286
# Unit Clause-clause subsumption calls : 457
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 19
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 467835

# -------------------------------------------------
# User time                : 0.360 s
# System time              : 0.007 s
# Total time               : 0.367 s
# Maximum resident set size: 28656 pages
