# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2v
# and selection function SelectCQArNTNp.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_waybel_7, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v1_subset_1(X2,u1_struct_0(X1))<=>~(r2_hidden(k3_yellow_0(X1),X2))))), file('data/MPTP/proofs/1b1edd66c44d4a6bbad0f6fd0a0bc148.E_input', t8_waybel_7)).
fof(t49_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(![X3]:(m1_subset_1(X3,X1)=>r2_hidden(X3,X2))=>X1=X2)), file('data/MPTP/proofs/1b1edd66c44d4a6bbad0f6fd0a0bc148.E_input', t49_subset_1)).
fof(t34_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:((r1_tarski(X2,X3)&(r1_yellow_0(X1,X2)&r1_yellow_0(X1,X3)))=>r1_orders_2(X1,k1_yellow_0(X1,X2),k1_yellow_0(X1,X3)))), file('data/MPTP/proofs/1b1edd66c44d4a6bbad0f6fd0a0bc148.E_input', t34_yellow_0)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/1b1edd66c44d4a6bbad0f6fd0a0bc148.E_input', t2_xboole_1)).
fof(t34_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r1_yellow_0(X1,k5_waybel_0(X1,X2))&k1_yellow_0(X1,k5_waybel_0(X1,X2))=X2))), file('data/MPTP/proofs/1b1edd66c44d4a6bbad0f6fd0a0bc148.E_input', t34_waybel_0)).
fof(t42_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))=>(r1_yellow_0(X1,k1_xboole_0)&r2_yellow_0(X1,u1_struct_0(X1)))), file('data/MPTP/proofs/1b1edd66c44d4a6bbad0f6fd0a0bc148.E_input', t42_yellow_0)).
fof(d20_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v13_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r1_orders_2(X1,X3,X4))=>r2_hidden(X4,X2))))))), file('data/MPTP/proofs/1b1edd66c44d4a6bbad0f6fd0a0bc148.E_input', d20_waybel_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/1b1edd66c44d4a6bbad0f6fd0a0bc148.E_input', t4_subset)).
fof(d11_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0)), file('data/MPTP/proofs/1b1edd66c44d4a6bbad0f6fd0a0bc148.E_input', d11_yellow_0)).
fof(d7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(v1_subset_1(X2,X1)<=>~(X2=X1))), file('data/MPTP/proofs/1b1edd66c44d4a6bbad0f6fd0a0bc148.E_input', d7_subset_1)).
fof(rc3_subset_1, axiom, ![X1]:?[X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))&~(v1_subset_1(X2,X1))), file('data/MPTP/proofs/1b1edd66c44d4a6bbad0f6fd0a0bc148.E_input', rc3_subset_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/1b1edd66c44d4a6bbad0f6fd0a0bc148.E_input', t2_subset)).
fof(dt_k3_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(k3_yellow_0(X1),u1_struct_0(X1))), file('data/MPTP/proofs/1b1edd66c44d4a6bbad0f6fd0a0bc148.E_input', dt_k3_yellow_0)).
fof(c_0_13, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v1_subset_1(X2,u1_struct_0(X1))<=>~(r2_hidden(k3_yellow_0(X1),X2)))))), inference(assume_negation,[status(cth)],[t8_waybel_7])).
fof(c_0_14, plain, ![X7, X8]:((m1_subset_1(esk3_2(X7,X8),X7)|X7=X8|~m1_subset_1(X8,k1_zfmisc_1(X7)))&(~r2_hidden(esk3_2(X7,X8),X8)|X7=X8|~m1_subset_1(X8,k1_zfmisc_1(X7)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t49_subset_1])])])])).
fof(c_0_15, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_yellow_0(esk1_0)&l1_orders_2(esk1_0))))))&((~v1_xboole_0(esk2_0)&(v2_waybel_0(esk2_0,esk1_0)&(v13_waybel_0(esk2_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))))&((~v1_subset_1(esk2_0,u1_struct_0(esk1_0))|r2_hidden(k3_yellow_0(esk1_0),esk2_0))&(v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k3_yellow_0(esk1_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_13])])])])).
fof(c_0_16, plain, ![X10, X11, X12]:(~l1_orders_2(X10)|(~r1_tarski(X11,X12)|(~r1_yellow_0(X10,X11)|~r1_yellow_0(X10,X12))|r1_orders_2(X10,k1_yellow_0(X10,X11),k1_yellow_0(X10,X12)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t34_yellow_0])])])).
fof(c_0_17, plain, ![X13]:r1_tarski(k1_xboole_0,X13), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_18, plain, ![X14, X15]:((r1_yellow_0(X14,k5_waybel_0(X14,X15))|~m1_subset_1(X15,u1_struct_0(X14))|(v2_struct_0(X14)|(~v3_orders_2(X14)|(~v4_orders_2(X14)|(~v5_orders_2(X14)|~l1_orders_2(X14))))))&(k1_yellow_0(X14,k5_waybel_0(X14,X15))=X15|~m1_subset_1(X15,u1_struct_0(X14))|(v2_struct_0(X14)|(~v3_orders_2(X14)|(~v4_orders_2(X14)|(~v5_orders_2(X14)|~l1_orders_2(X14))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_waybel_0])])])])])).
cnf(c_0_19, plain, (m1_subset_1(esk3_2(X1,X2),X1)|X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_21, plain, ![X16]:((r1_yellow_0(X16,k1_xboole_0)|(v2_struct_0(X16)|(~v5_orders_2(X16)|(~v1_yellow_0(X16)|~l1_orders_2(X16)))))&(r2_yellow_0(X16,u1_struct_0(X16))|(v2_struct_0(X16)|(~v5_orders_2(X16)|(~v1_yellow_0(X16)|~l1_orders_2(X16)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t42_yellow_0])])])])).
cnf(c_0_22, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),k1_yellow_0(X1,X3))|~l1_orders_2(X1)|~r1_tarski(X2,X3)|~r1_yellow_0(X1,X2)|~r1_yellow_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (r1_yellow_0(X1,k5_waybel_0(X1,X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|m1_subset_1(esk3_2(u1_struct_0(esk1_0),esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_26, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_27, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_28, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_29, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_30, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_31, plain, (r1_yellow_0(X1,k1_xboole_0)|v2_struct_0(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_32, negated_conjecture, (v1_yellow_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_33, plain, (r1_orders_2(X1,k1_yellow_0(X1,k1_xboole_0),k1_yellow_0(X1,X2))|~r1_yellow_0(X1,k1_xboole_0)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_34, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r1_yellow_0(esk1_0,k5_waybel_0(esk1_0,esk3_2(u1_struct_0(esk1_0),esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29])]), c_0_30])).
cnf(c_0_35, negated_conjecture, (r1_yellow_0(esk1_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_26]), c_0_27])]), c_0_30])).
cnf(c_0_36, plain, (k1_yellow_0(X1,k5_waybel_0(X1,X2))=X2|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_37, plain, ![X17, X18, X19, X20]:((~v13_waybel_0(X18,X17)|(~m1_subset_1(X19,u1_struct_0(X17))|(~m1_subset_1(X20,u1_struct_0(X17))|(~r2_hidden(X19,X18)|~r1_orders_2(X17,X19,X20)|r2_hidden(X20,X18))))|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))|~l1_orders_2(X17))&((m1_subset_1(esk4_2(X17,X18),u1_struct_0(X17))|v13_waybel_0(X18,X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))|~l1_orders_2(X17))&((m1_subset_1(esk5_2(X17,X18),u1_struct_0(X17))|v13_waybel_0(X18,X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))|~l1_orders_2(X17))&(((r2_hidden(esk4_2(X17,X18),X18)|v13_waybel_0(X18,X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))|~l1_orders_2(X17))&(r1_orders_2(X17,esk4_2(X17,X18),esk5_2(X17,X18))|v13_waybel_0(X18,X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))|~l1_orders_2(X17)))&(~r2_hidden(esk5_2(X17,X18),X18)|v13_waybel_0(X18,X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))|~l1_orders_2(X17)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d20_waybel_0])])])])])).
fof(c_0_38, plain, ![X23, X24, X25]:(~r2_hidden(X23,X24)|~m1_subset_1(X24,k1_zfmisc_1(X25))|m1_subset_1(X23,X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_39, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,k1_xboole_0),k1_yellow_0(esk1_0,k5_waybel_0(esk1_0,esk3_2(u1_struct_0(esk1_0),esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35]), c_0_26])])).
cnf(c_0_40, negated_conjecture, (k1_yellow_0(esk1_0,k5_waybel_0(esk1_0,esk3_2(u1_struct_0(esk1_0),esk2_0)))=esk3_2(u1_struct_0(esk1_0),esk2_0)|u1_struct_0(esk1_0)=esk2_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29])]), c_0_30])).
fof(c_0_41, plain, ![X26]:(~l1_orders_2(X26)|k3_yellow_0(X26)=k1_yellow_0(X26,k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_yellow_0])])).
fof(c_0_42, plain, ![X27, X28]:((~v1_subset_1(X28,X27)|X28!=X27|~m1_subset_1(X28,k1_zfmisc_1(X27)))&(X28=X27|v1_subset_1(X28,X27)|~m1_subset_1(X28,k1_zfmisc_1(X27)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_subset_1])])])])).
cnf(c_0_43, plain, (r2_hidden(X4,X1)|~v13_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,X1)|~r1_orders_2(X2,X3,X4)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_44, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_45, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,k1_xboole_0),esk3_2(u1_struct_0(esk1_0),esk2_0))), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_46, plain, (k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_47, plain, (X1=X2|v1_subset_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_48, plain, ![X29]:(m1_subset_1(esk6_1(X29),k1_zfmisc_1(X29))&~v1_subset_1(esk6_1(X29),X29)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[rc3_subset_1])])])).
cnf(c_0_49, plain, (r2_hidden(X1,X2)|~r1_orders_2(X3,X4,X1)|~r2_hidden(X4,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~v13_waybel_0(X2,X3)|~l1_orders_2(X3)), inference(csr,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_50, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r1_orders_2(esk1_0,k3_yellow_0(esk1_0),esk3_2(u1_struct_0(esk1_0),esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_26])])).
cnf(c_0_51, negated_conjecture, (r2_hidden(k3_yellow_0(esk1_0),esk2_0)|~v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_52, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_47, c_0_20])).
cnf(c_0_53, plain, (X1=X2|~r2_hidden(esk3_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_54, plain, (m1_subset_1(esk6_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_55, plain, (~v1_subset_1(esk6_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_56, plain, ![X31, X32]:(~m1_subset_1(X31,X32)|(v1_xboole_0(X32)|r2_hidden(X31,X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_57, plain, ![X33]:(~l1_orders_2(X33)|m1_subset_1(k3_yellow_0(X33),u1_struct_0(X33))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_yellow_0])])).
cnf(c_0_58, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_hidden(esk3_2(u1_struct_0(esk1_0),esk2_0),X1)|~r2_hidden(k3_yellow_0(esk1_0),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v13_waybel_0(X1,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_26])]), c_0_25])).
cnf(c_0_59, negated_conjecture, (v13_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_60, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_hidden(k3_yellow_0(esk1_0),esk2_0)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_61, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|~r2_hidden(esk3_2(u1_struct_0(esk1_0),esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_53, c_0_20])).
cnf(c_0_62, plain, (esk6_1(X1)=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_54]), c_0_55])).
cnf(c_0_63, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_64, plain, (m1_subset_1(k3_yellow_0(X1),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_65, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k3_yellow_0(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_66, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_20]), c_0_59])]), c_0_60]), c_0_61])).
cnf(c_0_67, plain, (~v1_subset_1(X1,X1)), inference(rw,[status(thm)],[c_0_55, c_0_62])).
cnf(c_0_68, plain, (r2_hidden(k3_yellow_0(X1),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_69, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_70, negated_conjecture, (~r2_hidden(k3_yellow_0(esk1_0),esk2_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_66]), c_0_67])).
cnf(c_0_71, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_66]), c_0_26])]), c_0_69]), c_0_70]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 72
# Proof object clause steps            : 45
# Proof object formula steps           : 27
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 37
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 34
# Processed clauses                    : 131
# ...of these trivial                  : 1
# ...subsumed                          : 11
# ...remaining for further processing  : 118
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 30
# Generated clauses                    : 100
# ...of the previous two non-trivial   : 85
# Contextual simplify-reflections      : 7
# Paramodulations                      : 99
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 53
#    Positive orientable unit clauses  : 16
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 33
# Current number of unprocessed clauses: 21
# ...number of literals in the above   : 91
# Current number of archived formulas  : 0
# Current number of archived clauses   : 64
# Clause-clause subsumption calls (NU) : 601
# Rec. Clause-clause subsumption calls : 194
# Non-unit clause-clause subsumptions  : 17
# Unit Clause-clause subsumption calls : 16
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4651

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.003 s
# Total time               : 0.013 s
# Maximum resident set size: 28732 pages
