# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d9_relat_2, axiom, ![X1]:(v1_relat_1(X1)=>(v1_relat_2(X1)<=>r1_relat_2(X1,k3_relat_1(X1)))), file('data/MPTP/proofs/1b5f39000cdb45c5a8860d0c232a019b.E_input', d9_relat_2)).
fof(d5_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(r2_wellord1(X1,X2)<=>(r1_relat_2(X1,X2)&(r8_relat_2(X1,X2)&(r4_relat_2(X1,X2)&(r6_relat_2(X1,X2)&r1_wellord1(X1,X2))))))), file('data/MPTP/proofs/1b5f39000cdb45c5a8860d0c232a019b.E_input', d5_wellord1)).
fof(t8_wellord1, conjecture, ![X1]:(v1_relat_1(X1)=>(r2_wellord1(X1,k3_relat_1(X1))<=>v2_wellord1(X1))), file('data/MPTP/proofs/1b5f39000cdb45c5a8860d0c232a019b.E_input', t8_wellord1)).
fof(d12_relat_2, axiom, ![X1]:(v1_relat_1(X1)=>(v4_relat_2(X1)<=>r4_relat_2(X1,k3_relat_1(X1)))), file('data/MPTP/proofs/1b5f39000cdb45c5a8860d0c232a019b.E_input', d12_relat_2)).
fof(d16_relat_2, axiom, ![X1]:(v1_relat_1(X1)=>(v8_relat_2(X1)<=>r8_relat_2(X1,k3_relat_1(X1)))), file('data/MPTP/proofs/1b5f39000cdb45c5a8860d0c232a019b.E_input', d16_relat_2)).
fof(t5_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v1_wellord1(X1)<=>r1_wellord1(X1,k3_relat_1(X1)))), file('data/MPTP/proofs/1b5f39000cdb45c5a8860d0c232a019b.E_input', t5_wellord1)).
fof(d14_relat_2, axiom, ![X1]:(v1_relat_1(X1)=>(v6_relat_2(X1)<=>r6_relat_2(X1,k3_relat_1(X1)))), file('data/MPTP/proofs/1b5f39000cdb45c5a8860d0c232a019b.E_input', d14_relat_2)).
fof(d4_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v2_wellord1(X1)<=>(v1_relat_2(X1)&(v8_relat_2(X1)&(v4_relat_2(X1)&(v6_relat_2(X1)&v1_wellord1(X1))))))), file('data/MPTP/proofs/1b5f39000cdb45c5a8860d0c232a019b.E_input', d4_wellord1)).
fof(c_0_8, plain, ![X10]:((~v1_relat_2(X10)|r1_relat_2(X10,k3_relat_1(X10))|~v1_relat_1(X10))&(~r1_relat_2(X10,k3_relat_1(X10))|v1_relat_2(X10)|~v1_relat_1(X10))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_relat_2])])])).
fof(c_0_9, plain, ![X3, X4]:(((r1_relat_2(X3,X4)|~r2_wellord1(X3,X4)|~v1_relat_1(X3))&((r8_relat_2(X3,X4)|~r2_wellord1(X3,X4)|~v1_relat_1(X3))&((r4_relat_2(X3,X4)|~r2_wellord1(X3,X4)|~v1_relat_1(X3))&((r6_relat_2(X3,X4)|~r2_wellord1(X3,X4)|~v1_relat_1(X3))&(r1_wellord1(X3,X4)|~r2_wellord1(X3,X4)|~v1_relat_1(X3))))))&(~r1_relat_2(X3,X4)|(~r8_relat_2(X3,X4)|(~r4_relat_2(X3,X4)|(~r6_relat_2(X3,X4)|~r1_wellord1(X3,X4))))|r2_wellord1(X3,X4)|~v1_relat_1(X3))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_wellord1])])])])).
fof(c_0_10, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>(r2_wellord1(X1,k3_relat_1(X1))<=>v2_wellord1(X1)))), inference(assume_negation,[status(cth)],[t8_wellord1])).
fof(c_0_11, plain, ![X7]:((~v4_relat_2(X7)|r4_relat_2(X7,k3_relat_1(X7))|~v1_relat_1(X7))&(~r4_relat_2(X7,k3_relat_1(X7))|v4_relat_2(X7)|~v1_relat_1(X7))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_relat_2])])])).
fof(c_0_12, plain, ![X8]:((~v8_relat_2(X8)|r8_relat_2(X8,k3_relat_1(X8))|~v1_relat_1(X8))&(~r8_relat_2(X8,k3_relat_1(X8))|v8_relat_2(X8)|~v1_relat_1(X8))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d16_relat_2])])])).
fof(c_0_13, plain, ![X9]:((~v1_wellord1(X9)|r1_wellord1(X9,k3_relat_1(X9))|~v1_relat_1(X9))&(~r1_wellord1(X9,k3_relat_1(X9))|v1_wellord1(X9)|~v1_relat_1(X9))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_wellord1])])])).
cnf(c_0_14, plain, (v1_relat_2(X1)|~r1_relat_2(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15, plain, (r1_relat_2(X1,X2)|~r2_wellord1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_16, negated_conjecture, (v1_relat_1(esk1_0)&((~r2_wellord1(esk1_0,k3_relat_1(esk1_0))|~v2_wellord1(esk1_0))&(r2_wellord1(esk1_0,k3_relat_1(esk1_0))|v2_wellord1(esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
cnf(c_0_17, plain, (v4_relat_2(X1)|~r4_relat_2(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, plain, (r4_relat_2(X1,X2)|~r2_wellord1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_19, plain, (v8_relat_2(X1)|~r8_relat_2(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20, plain, (r8_relat_2(X1,X2)|~r2_wellord1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_21, plain, ![X5]:((~v6_relat_2(X5)|r6_relat_2(X5,k3_relat_1(X5))|~v1_relat_1(X5))&(~r6_relat_2(X5,k3_relat_1(X5))|v6_relat_2(X5)|~v1_relat_1(X5))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d14_relat_2])])])).
cnf(c_0_22, plain, (r2_wellord1(X1,X2)|~r1_relat_2(X1,X2)|~r8_relat_2(X1,X2)|~r4_relat_2(X1,X2)|~r6_relat_2(X1,X2)|~r1_wellord1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_23, plain, (r1_wellord1(X1,k3_relat_1(X1))|~v1_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_24, plain, ![X11]:(((v1_relat_2(X11)|~v2_wellord1(X11)|~v1_relat_1(X11))&((v8_relat_2(X11)|~v2_wellord1(X11)|~v1_relat_1(X11))&((v4_relat_2(X11)|~v2_wellord1(X11)|~v1_relat_1(X11))&((v6_relat_2(X11)|~v2_wellord1(X11)|~v1_relat_1(X11))&(v1_wellord1(X11)|~v2_wellord1(X11)|~v1_relat_1(X11))))))&(~v1_relat_2(X11)|(~v8_relat_2(X11)|(~v4_relat_2(X11)|(~v6_relat_2(X11)|~v1_wellord1(X11))))|v2_wellord1(X11)|~v1_relat_1(X11))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_wellord1])])])).
cnf(c_0_25, plain, (v1_relat_2(X1)|~r2_wellord1(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_26, negated_conjecture, (r2_wellord1(esk1_0,k3_relat_1(esk1_0))|v2_wellord1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_27, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28, plain, (v4_relat_2(X1)|~r2_wellord1(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_29, plain, (v8_relat_2(X1)|~r2_wellord1(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_30, plain, (v1_wellord1(X1)|~r1_wellord1(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_31, plain, (r1_wellord1(X1,X2)|~r2_wellord1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_32, plain, (v6_relat_2(X1)|~r6_relat_2(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_33, plain, (r6_relat_2(X1,X2)|~r2_wellord1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_34, plain, (r2_wellord1(X1,k3_relat_1(X1))|~v1_wellord1(X1)|~r6_relat_2(X1,k3_relat_1(X1))|~r4_relat_2(X1,k3_relat_1(X1))|~r8_relat_2(X1,k3_relat_1(X1))|~r1_relat_2(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_35, plain, (r6_relat_2(X1,k3_relat_1(X1))|~v6_relat_2(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_36, plain, (v2_wellord1(X1)|~v1_relat_2(X1)|~v8_relat_2(X1)|~v4_relat_2(X1)|~v6_relat_2(X1)|~v1_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_37, negated_conjecture, (v1_relat_2(esk1_0)|v2_wellord1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])])).
cnf(c_0_38, negated_conjecture, (v4_relat_2(esk1_0)|v2_wellord1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_26]), c_0_27])])).
cnf(c_0_39, negated_conjecture, (v8_relat_2(esk1_0)|v2_wellord1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_26]), c_0_27])])).
cnf(c_0_40, plain, (v1_wellord1(X1)|~r2_wellord1(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_41, plain, (v6_relat_2(X1)|~r2_wellord1(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_42, plain, (r2_wellord1(X1,k3_relat_1(X1))|~v1_wellord1(X1)|~v6_relat_2(X1)|~r4_relat_2(X1,k3_relat_1(X1))|~r8_relat_2(X1,k3_relat_1(X1))|~r1_relat_2(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_43, plain, (r4_relat_2(X1,k3_relat_1(X1))|~v4_relat_2(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_44, negated_conjecture, (v2_wellord1(esk1_0)|~v1_wellord1(esk1_0)|~v6_relat_2(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_27])]), c_0_38]), c_0_39])).
cnf(c_0_45, negated_conjecture, (v1_wellord1(esk1_0)|v2_wellord1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_26]), c_0_27])])).
cnf(c_0_46, negated_conjecture, (v2_wellord1(esk1_0)|v6_relat_2(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_26]), c_0_27])])).
cnf(c_0_47, plain, (r2_wellord1(X1,k3_relat_1(X1))|~v1_wellord1(X1)|~v4_relat_2(X1)|~v6_relat_2(X1)|~r8_relat_2(X1,k3_relat_1(X1))|~r1_relat_2(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_48, plain, (r8_relat_2(X1,k3_relat_1(X1))|~v8_relat_2(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_49, negated_conjecture, (~r2_wellord1(esk1_0,k3_relat_1(esk1_0))|~v2_wellord1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_50, negated_conjecture, (v2_wellord1(esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])).
cnf(c_0_51, plain, (r2_wellord1(X1,k3_relat_1(X1))|~v1_wellord1(X1)|~v8_relat_2(X1)|~v4_relat_2(X1)|~v6_relat_2(X1)|~r1_relat_2(X1,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_52, plain, (r1_relat_2(X1,k3_relat_1(X1))|~v1_relat_2(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_53, negated_conjecture, (~r2_wellord1(esk1_0,k3_relat_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_50])])).
cnf(c_0_54, plain, (r2_wellord1(X1,k3_relat_1(X1))|~v1_relat_2(X1)|~v1_wellord1(X1)|~v8_relat_2(X1)|~v4_relat_2(X1)|~v6_relat_2(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_55, negated_conjecture, (~v1_relat_2(esk1_0)|~v1_wellord1(esk1_0)|~v8_relat_2(esk1_0)|~v4_relat_2(esk1_0)|~v6_relat_2(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_27])])).
cnf(c_0_56, plain, (v1_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_57, negated_conjecture, (~v1_wellord1(esk1_0)|~v8_relat_2(esk1_0)|~v4_relat_2(esk1_0)|~v6_relat_2(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_50]), c_0_27])])).
cnf(c_0_58, plain, (v1_wellord1(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_59, negated_conjecture, (~v8_relat_2(esk1_0)|~v4_relat_2(esk1_0)|~v6_relat_2(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_50]), c_0_27])])).
cnf(c_0_60, plain, (v8_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_61, negated_conjecture, (~v4_relat_2(esk1_0)|~v6_relat_2(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_50]), c_0_27])])).
cnf(c_0_62, plain, (v4_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_63, negated_conjecture, (~v6_relat_2(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_50]), c_0_27])])).
cnf(c_0_64, plain, (v6_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_65, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_50]), c_0_27])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 66
# Proof object clause steps            : 49
# Proof object formula steps           : 17
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 23
# Proof object simplifying inferences  : 34
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 25
# Processed clauses                    : 73
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 73
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 8
# Generated clauses                    : 41
# ...of the previous two non-trivial   : 23
# Contextual simplify-reflections      : 3
# Paramodulations                      : 41
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 37
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 33
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 36
# Clause-clause subsumption calls (NU) : 373
# Rec. Clause-clause subsumption calls : 135
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2398

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28764 pages
