# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/1bccdd2921944c48947474d2c208506c.E_input', d1_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/1bccdd2921944c48947474d2c208506c.E_input', d3_tarski)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/1bccdd2921944c48947474d2c208506c.E_input', cc2_relat_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/1bccdd2921944c48947474d2c208506c.E_input', t3_subset)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/1bccdd2921944c48947474d2c208506c.E_input', d10_xboole_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/1bccdd2921944c48947474d2c208506c.E_input', fc1_xboole_0)).
fof(t25_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>(r1_tarski(k1_relat_1(X1),k1_relat_1(X2))&r1_tarski(k2_relat_1(X1),k2_relat_1(X2)))))), file('data/MPTP/proofs/1bccdd2921944c48947474d2c208506c.E_input', t25_relat_1)).
fof(t194_relat_1, axiom, ![X1, X2]:r1_tarski(k2_relat_1(k2_zfmisc_1(X1,X2)),X2), file('data/MPTP/proofs/1bccdd2921944c48947474d2c208506c.E_input', t194_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/1bccdd2921944c48947474d2c208506c.E_input', fc6_relat_1)).
fof(t4_funct_2, conjecture, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r1_tarski(k2_relat_1(X2),X1)=>(v1_funct_1(X2)&(v1_funct_2(X2,k1_relat_1(X2),X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X2),X1))))))), file('data/MPTP/proofs/1bccdd2921944c48947474d2c208506c.E_input', t4_funct_2)).
fof(t65_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>(k1_relat_1(X1)=k1_xboole_0<=>k2_relat_1(X1)=k1_xboole_0)), file('data/MPTP/proofs/1bccdd2921944c48947474d2c208506c.E_input', t65_relat_1)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/1bccdd2921944c48947474d2c208506c.E_input', t1_xboole_1)).
fof(t193_relat_1, axiom, ![X1, X2]:r1_tarski(k1_relat_1(k2_zfmisc_1(X1,X2)),X1), file('data/MPTP/proofs/1bccdd2921944c48947474d2c208506c.E_input', t193_relat_1)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(((X2=k1_xboole_0=>X1=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X2,X3)))&(X2=k1_xboole_0=>(X1=k1_xboole_0|(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0))))), file('data/MPTP/proofs/1bccdd2921944c48947474d2c208506c.E_input', d1_funct_2)).
fof(t11_relset_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>((r1_tarski(k1_relat_1(X3),X1)&r1_tarski(k2_relat_1(X3),X2))=>m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))), file('data/MPTP/proofs/1bccdd2921944c48947474d2c208506c.E_input', t11_relset_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/1bccdd2921944c48947474d2c208506c.E_input', redefinition_k1_relset_1)).
fof(c_0_16, plain, ![X22, X23, X24]:((~v1_xboole_0(X22)|~r2_hidden(X23,X22))&(r2_hidden(esk4_1(X24),X24)|v1_xboole_0(X24))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_17, plain, ![X11, X12, X13, X14, X15]:((~r1_tarski(X11,X12)|(~r2_hidden(X13,X11)|r2_hidden(X13,X12)))&((r2_hidden(esk3_2(X14,X15),X14)|r1_tarski(X14,X15))&(~r2_hidden(esk3_2(X14,X15),X15)|r1_tarski(X14,X15)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_18, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_19, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_20, plain, ![X41, X42]:(~v1_relat_1(X41)|(~m1_subset_1(X42,k1_zfmisc_1(X41))|v1_relat_1(X42))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_21, plain, ![X39, X40]:((~m1_subset_1(X39,k1_zfmisc_1(X40))|r1_tarski(X39,X40))&(~r1_tarski(X39,X40)|m1_subset_1(X39,k1_zfmisc_1(X40)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_22, plain, ![X37, X38]:(((r1_tarski(X37,X38)|X37!=X38)&(r1_tarski(X38,X37)|X37!=X38))&(~r1_tarski(X37,X38)|~r1_tarski(X38,X37)|X37=X38)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_23, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_24, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
fof(c_0_25, plain, ![X20, X21]:((r1_tarski(k1_relat_1(X20),k1_relat_1(X21))|~r1_tarski(X20,X21)|~v1_relat_1(X21)|~v1_relat_1(X20))&(r1_tarski(k2_relat_1(X20),k2_relat_1(X21))|~r1_tarski(X20,X21)|~v1_relat_1(X21)|~v1_relat_1(X20))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t25_relat_1])])])])).
cnf(c_0_26, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_30, plain, ![X9, X10]:r1_tarski(k2_relat_1(k2_zfmisc_1(X9,X10)),X10), inference(variable_rename,[status(thm)],[t194_relat_1])).
cnf(c_0_31, plain, (r1_tarski(k2_relat_1(X1),k2_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, plain, (v1_relat_1(X1)|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_33, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34, plain, (r1_tarski(k2_relat_1(k2_zfmisc_1(X1,X2)),X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_35, plain, ![X29, X30]:v1_relat_1(k2_zfmisc_1(X29,X30)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_36, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_37, plain, (r1_tarski(k2_relat_1(X1),k2_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(csr,[status(thm)],[c_0_31, c_0_32])).
fof(c_0_38, negated_conjecture, ~(![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r1_tarski(k2_relat_1(X2),X1)=>(v1_funct_1(X2)&(v1_funct_2(X2,k1_relat_1(X2),X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X2),X1)))))))), inference(assume_negation,[status(cth)],[t4_funct_2])).
fof(c_0_39, plain, ![X31]:((k1_relat_1(X31)!=k1_xboole_0|k2_relat_1(X31)=k1_xboole_0|~v1_relat_1(X31))&(k2_relat_1(X31)!=k1_xboole_0|k1_relat_1(X31)=k1_xboole_0|~v1_relat_1(X31))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t65_relat_1])])])).
cnf(c_0_40, plain, (k2_relat_1(k2_zfmisc_1(X1,k1_xboole_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_41, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_42, plain, (r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(X1,k2_relat_1(X3))|~r1_tarski(X3,X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_43, plain, (r2_hidden(esk4_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_44, plain, ![X17, X18, X19]:(~r1_tarski(X17,X18)|~r1_tarski(X18,X19)|r1_tarski(X17,X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_45, negated_conjecture, ((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&(r1_tarski(k2_relat_1(esk2_0),esk1_0)&(~v1_funct_1(esk2_0)|(~v1_funct_2(esk2_0,k1_relat_1(esk2_0),esk1_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(esk2_0),esk1_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_38])])])).
fof(c_0_46, plain, ![X32, X33]:r1_tarski(k1_relat_1(k2_zfmisc_1(X32,X33)),X32), inference(variable_rename,[status(thm)],[t193_relat_1])).
cnf(c_0_47, plain, (k2_relat_1(X1)=k1_xboole_0|k1_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_48, plain, (r1_tarski(k2_relat_1(X1),k1_xboole_0)|~r1_tarski(X1,k2_zfmisc_1(X2,k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_40]), c_0_41])])).
cnf(c_0_49, plain, (v1_xboole_0(k2_relat_1(X1))|r2_hidden(esk4_1(k2_relat_1(X1)),k2_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_50, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_51, negated_conjecture, (r1_tarski(k2_relat_1(esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_52, plain, (r1_tarski(k1_relat_1(X1),k1_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_53, plain, (r1_tarski(k1_relat_1(k2_zfmisc_1(X1,X2)),X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_54, plain, (r1_tarski(k2_relat_1(X1),k1_xboole_0)|k1_relat_1(X2)!=k1_xboole_0|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_37, c_0_47])).
cnf(c_0_55, plain, (r1_tarski(k2_relat_1(k1_xboole_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_48, c_0_29])).
cnf(c_0_56, plain, (v1_relat_1(k1_xboole_0)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_32, c_0_29])).
cnf(c_0_57, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_58, plain, (v1_xboole_0(k2_relat_1(X1))|~v1_xboole_0(k2_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_18, c_0_49])).
cnf(c_0_59, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_60, negated_conjecture, (r1_tarski(X1,esk1_0)|~r1_tarski(X1,k2_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_61, plain, (r1_tarski(k1_relat_1(X1),k1_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(csr,[status(thm)],[c_0_52, c_0_32])).
cnf(c_0_62, plain, (k1_relat_1(k2_zfmisc_1(k1_xboole_0,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_33, c_0_53])).
cnf(c_0_63, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_50, c_0_29])).
cnf(c_0_64, negated_conjecture, (r1_tarski(k2_relat_1(k2_relat_1(esk2_0)),k1_xboole_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_54, c_0_51])).
cnf(c_0_65, plain, (k1_relat_1(X1)=k1_xboole_0|k2_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_66, plain, (k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_55]), c_0_29])])).
cnf(c_0_67, negated_conjecture, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_68, plain, (v1_xboole_0(k2_relat_1(X1))|k1_relat_1(X2)!=k1_xboole_0|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_47]), c_0_24])])).
cnf(c_0_69, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_59])).
cnf(c_0_70, negated_conjecture, (r1_tarski(k2_relat_1(k2_zfmisc_1(X1,k2_relat_1(esk2_0))),esk1_0)), inference(spm,[status(thm)],[c_0_60, c_0_34])).
cnf(c_0_71, plain, (r1_tarski(k1_relat_1(X1),k1_xboole_0)|~r1_tarski(X1,k2_zfmisc_1(k1_xboole_0,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_41])])).
cnf(c_0_72, negated_conjecture, (r1_tarski(k2_relat_1(k2_relat_1(esk2_0)),X1)|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_73, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67])])).
cnf(c_0_74, plain, (v1_xboole_0(k2_relat_1(X1))|k1_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_75, negated_conjecture, (r1_tarski(X1,esk1_0)|~r1_tarski(X1,k2_relat_1(k2_zfmisc_1(X2,k2_relat_1(esk2_0))))), inference(spm,[status(thm)],[c_0_50, c_0_70])).
cnf(c_0_76, negated_conjecture, (r1_tarski(k1_relat_1(k2_relat_1(k2_relat_1(esk2_0))),k1_xboole_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_77, plain, (r1_tarski(k1_relat_1(X1),k1_xboole_0)|~r1_tarski(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_73]), c_0_67])])).
cnf(c_0_78, plain, (r1_tarski(k2_relat_1(X1),k1_xboole_0)|~r1_tarski(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_66]), c_0_67])])).
cnf(c_0_79, plain, (k2_relat_1(X1)=k2_relat_1(X2)|~r1_tarski(k2_relat_1(X1),k2_relat_1(X2))|~r1_tarski(X2,X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_28, c_0_37])).
cnf(c_0_80, plain, (r1_tarski(k2_relat_1(X1),X2)|k1_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_23, c_0_74])).
cnf(c_0_81, negated_conjecture, (r1_tarski(k2_relat_1(X1),esk1_0)|~r1_tarski(X1,k2_zfmisc_1(X2,k2_relat_1(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_37]), c_0_41])])).
cnf(c_0_82, negated_conjecture, (r1_tarski(k1_relat_1(k2_relat_1(k2_relat_1(esk2_0))),X1)|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_63, c_0_76])).
cnf(c_0_83, negated_conjecture, (r1_tarski(k1_relat_1(k2_zfmisc_1(k2_relat_1(esk2_0),X1)),esk1_0)), inference(spm,[status(thm)],[c_0_60, c_0_53])).
cnf(c_0_84, plain, (r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_63, c_0_77])).
cnf(c_0_85, plain, (r1_tarski(k2_relat_1(X1),X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_63, c_0_78])).
cnf(c_0_86, plain, (k2_relat_1(X1)=k2_relat_1(X2)|k1_relat_1(X1)!=k1_xboole_0|~r1_tarski(X2,X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_87, negated_conjecture, (r1_tarski(k2_relat_1(k1_relat_1(k2_relat_1(k2_relat_1(esk2_0)))),esk1_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_88, negated_conjecture, (X1=k2_relat_1(k2_relat_1(esk2_0))|k1_relat_1(esk1_0)!=k1_xboole_0|~r1_tarski(X1,k2_relat_1(k2_relat_1(esk2_0)))|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_28, c_0_72])).
cnf(c_0_89, negated_conjecture, (r1_tarski(k2_relat_1(k2_relat_1(k2_relat_1(esk2_0))),k1_xboole_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_48, c_0_72])).
cnf(c_0_90, negated_conjecture, (r1_tarski(X1,esk1_0)|~r1_tarski(X1,k1_relat_1(k2_zfmisc_1(k2_relat_1(esk2_0),X2)))), inference(spm,[status(thm)],[c_0_50, c_0_83])).
cnf(c_0_91, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_relat_1(X3))|~r1_tarski(X3,k1_xboole_0)), inference(spm,[status(thm)],[c_0_36, c_0_84])).
cnf(c_0_92, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k2_relat_1(X3))|~r1_tarski(X3,k1_xboole_0)), inference(spm,[status(thm)],[c_0_50, c_0_85])).
cnf(c_0_93, negated_conjecture, (v1_relat_1(k2_relat_1(esk2_0))|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_32, c_0_51])).
cnf(c_0_94, negated_conjecture, (k2_relat_1(k2_relat_1(k1_relat_1(k2_relat_1(k2_relat_1(esk2_0)))))=k2_relat_1(esk1_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_95, negated_conjecture, (k1_relat_1(k2_relat_1(k2_relat_1(esk2_0)))=k2_relat_1(k2_relat_1(esk2_0))|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_88, c_0_82])).
cnf(c_0_96, negated_conjecture, (r1_tarski(k2_relat_1(k2_relat_1(k2_relat_1(esk2_0))),X1)|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_63, c_0_89])).
cnf(c_0_97, negated_conjecture, (r1_tarski(k1_relat_1(X1),esk1_0)|~r1_tarski(X1,k2_zfmisc_1(k2_relat_1(esk2_0),X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_61]), c_0_41])])).
cnf(c_0_98, plain, (v1_xboole_0(k1_relat_1(X1))|r2_hidden(esk4_1(k1_relat_1(X1)),X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_91, c_0_43])).
cnf(c_0_99, plain, (r1_tarski(k2_relat_1(k2_relat_1(X1)),k1_xboole_0)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_48, c_0_85])).
cnf(c_0_100, plain, (r1_tarski(k2_relat_1(X1),X2)|~r1_tarski(X3,k1_xboole_0)|~r1_tarski(X1,X3)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_92, c_0_37])).
cnf(c_0_101, plain, (v1_relat_1(k2_relat_1(X1))|~r1_tarski(X1,X2)|~v1_relat_1(k2_relat_1(X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_32, c_0_37])).
cnf(c_0_102, negated_conjecture, (k2_relat_1(k2_relat_1(esk2_0))=k2_relat_1(X1)|k1_relat_1(esk1_0)!=k1_xboole_0|~r1_tarski(X1,k2_relat_1(esk2_0))|~v1_relat_1(esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_72]), c_0_93])).
cnf(c_0_103, negated_conjecture, (k2_relat_1(k2_relat_1(k2_relat_1(k2_relat_1(esk2_0))))=k2_relat_1(esk1_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_104, negated_conjecture, (k2_relat_1(k2_relat_1(k2_relat_1(esk2_0)))=k2_relat_1(k2_relat_1(esk2_0))|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_88, c_0_96])).
cnf(c_0_105, negated_conjecture, (X1=k1_relat_1(k2_relat_1(k2_relat_1(esk2_0)))|k1_relat_1(esk1_0)!=k1_xboole_0|~r1_tarski(X1,k1_relat_1(k2_relat_1(k2_relat_1(esk2_0))))|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_28, c_0_82])).
cnf(c_0_106, negated_conjecture, (r1_tarski(k1_relat_1(k1_relat_1(k2_relat_1(k2_relat_1(esk2_0)))),esk1_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_97, c_0_82])).
cnf(c_0_107, negated_conjecture, (k2_relat_1(k2_relat_1(esk2_0))=k2_relat_1(esk1_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_86, c_0_51])).
cnf(c_0_108, plain, (v1_xboole_0(k1_relat_1(X1))|~v1_xboole_0(X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_18, c_0_98])).
cnf(c_0_109, plain, (X1=k1_relat_1(X2)|~r1_tarski(X1,k1_relat_1(X2))|~r1_tarski(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_28, c_0_84])).
cnf(c_0_110, plain, (r1_tarski(k2_relat_1(k2_relat_1(X1)),X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_63, c_0_99])).
cnf(c_0_111, negated_conjecture, (r1_tarski(k2_relat_1(k2_relat_1(esk2_0)),X1)|~r1_tarski(esk1_0,k1_xboole_0)|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_100, c_0_51])).
cnf(c_0_112, plain, (X1=k2_relat_1(X2)|~r1_tarski(X1,k2_relat_1(X2))|~r1_tarski(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_28, c_0_85])).
cnf(c_0_113, plain, (v1_relat_1(k2_relat_1(X1))|k1_relat_1(X2)!=k1_xboole_0|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_47]), c_0_67])])).
cnf(c_0_114, negated_conjecture, (k2_relat_1(k2_relat_1(k2_relat_1(k2_relat_1(esk2_0))))=k2_relat_1(k2_relat_1(esk2_0))|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_102, c_0_96])).
cnf(c_0_115, negated_conjecture, (k2_relat_1(k2_relat_1(k2_relat_1(esk2_0)))=k2_relat_1(esk1_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_116, negated_conjecture, (k2_relat_1(k2_relat_1(k2_relat_1(esk2_0)))=k1_relat_1(k2_relat_1(k2_relat_1(esk2_0)))|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_105, c_0_96])).
cnf(c_0_117, plain, (k1_relat_1(X1)=k1_relat_1(X2)|~r1_tarski(k1_relat_1(X1),k1_relat_1(X2))|~r1_tarski(X2,X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_28, c_0_61])).
cnf(c_0_118, negated_conjecture, (r1_tarski(k1_relat_1(k1_relat_1(k2_relat_1(esk1_0))),esk1_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_106, c_0_107])).
cnf(c_0_119, negated_conjecture, (k1_relat_1(k2_relat_1(esk1_0))=k2_relat_1(esk1_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_95, c_0_107])).
cnf(c_0_120, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X1,k2_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_36, c_0_51])).
cnf(c_0_121, plain, (v1_xboole_0(k1_relat_1(X1))|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_108, c_0_24])).
cnf(c_0_122, plain, (k2_relat_1(k2_relat_1(X1))=k1_relat_1(X2)|~r1_tarski(X2,k1_xboole_0)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_109, c_0_110])).
cnf(c_0_123, negated_conjecture, (r1_tarski(k2_relat_1(k2_relat_1(k2_relat_1(esk2_0))),esk1_0)|~r1_tarski(esk1_0,k1_xboole_0)|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_81, c_0_111])).
cnf(c_0_124, plain, (k2_relat_1(X1)=k2_relat_1(X2)|~r1_tarski(X2,k1_xboole_0)|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_112, c_0_37])).
cnf(c_0_125, negated_conjecture, (v1_relat_1(k2_relat_1(k2_relat_1(esk2_0)))|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_113, c_0_51])).
cnf(c_0_126, negated_conjecture, (k2_relat_1(k2_relat_1(esk1_0))=k2_relat_1(k2_relat_1(esk2_0))|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_127, negated_conjecture, (k2_relat_1(k2_relat_1(esk1_0))=k1_relat_1(k2_relat_1(esk1_0))|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_116, c_0_107])).
cnf(c_0_128, plain, (k1_relat_1(X1)=k1_relat_1(X2)|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_61]), c_0_32])).
cnf(c_0_129, negated_conjecture, (r1_tarski(k1_relat_1(k2_relat_1(esk1_0)),esk1_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_118, c_0_119])).
cnf(c_0_130, negated_conjecture, (v1_xboole_0(k2_relat_1(esk2_0))|r2_hidden(esk4_1(k2_relat_1(esk2_0)),esk1_0)), inference(spm,[status(thm)],[c_0_120, c_0_43])).
cnf(c_0_131, plain, (v1_xboole_0(k2_relat_1(k2_relat_1(X1)))|~r1_tarski(X2,k1_xboole_0)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_121, c_0_122])).
cnf(c_0_132, negated_conjecture, (k2_relat_1(k2_relat_1(k2_relat_1(esk2_0)))=esk1_0|~r1_tarski(esk1_0,k1_xboole_0)|~v1_relat_1(esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_123]), c_0_63])).
cnf(c_0_133, negated_conjecture, (k2_relat_1(k2_relat_1(esk2_0))=k2_relat_1(esk1_0)|~r1_tarski(esk1_0,k1_xboole_0)|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_124, c_0_51])).
cnf(c_0_134, negated_conjecture, (k1_relat_1(k2_relat_1(k2_relat_1(esk2_0)))=k1_relat_1(X1)|k1_relat_1(esk1_0)!=k1_xboole_0|~r1_tarski(X1,k2_relat_1(k2_relat_1(esk2_0)))|~v1_relat_1(esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_82]), c_0_125])).
cnf(c_0_135, negated_conjecture, (k2_relat_1(k1_relat_1(k2_relat_1(k2_relat_1(esk2_0))))=k2_relat_1(k2_relat_1(esk2_0))|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_102, c_0_82])).
cnf(c_0_136, negated_conjecture, (k2_relat_1(k2_relat_1(esk2_0))=k1_relat_1(k2_relat_1(esk1_0))|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_137, negated_conjecture, (k1_relat_1(k1_relat_1(k2_relat_1(esk1_0)))=k1_relat_1(esk1_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~r1_tarski(esk1_0,k1_relat_1(k2_relat_1(esk1_0)))|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_128, c_0_129])).
cnf(c_0_138, negated_conjecture, (k1_relat_1(k2_relat_1(k2_relat_1(esk2_0)))=k2_relat_1(esk1_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_115, c_0_116])).
cnf(c_0_139, negated_conjecture, (v1_xboole_0(k2_relat_1(esk2_0))|~v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_18, c_0_130])).
cnf(c_0_140, plain, (v1_xboole_0(k2_relat_1(k2_relat_1(X1)))|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_131, c_0_29])).
cnf(c_0_141, negated_conjecture, (k2_relat_1(k2_relat_1(esk1_0))=esk1_0|~r1_tarski(esk1_0,k1_xboole_0)|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_132, c_0_133])).
cnf(c_0_142, negated_conjecture, (k1_relat_1(k2_relat_1(k2_relat_1(k2_relat_1(esk2_0))))=k1_relat_1(k2_relat_1(k2_relat_1(esk2_0)))|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_134, c_0_96])).
cnf(c_0_143, negated_conjecture, (k2_relat_1(k1_relat_1(k1_relat_1(k2_relat_1(esk1_0))))=k1_relat_1(k2_relat_1(esk1_0))|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_144, negated_conjecture, (k1_relat_1(k2_relat_1(esk1_0))=k1_relat_1(esk1_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~r1_tarski(esk1_0,k2_relat_1(esk1_0))|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_137, c_0_119])).
cnf(c_0_145, negated_conjecture, (k1_relat_1(k1_relat_1(k2_relat_1(esk1_0)))=k2_relat_1(esk1_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_138, c_0_136])).
cnf(c_0_146, negated_conjecture, (r1_tarski(k2_relat_1(esk2_0),X1)|~v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_23, c_0_139])).
cnf(c_0_147, negated_conjecture, (v1_xboole_0(esk1_0)|~r1_tarski(esk1_0,k1_xboole_0)|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_140, c_0_141])).
fof(c_0_148, plain, ![X6, X7, X8]:((((~v1_funct_2(X8,X6,X7)|X6=k1_relset_1(X6,X7,X8)|X7=k1_xboole_0|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))&(X6!=k1_relset_1(X6,X7,X8)|v1_funct_2(X8,X6,X7)|X7=k1_xboole_0|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))))&((~v1_funct_2(X8,X6,X7)|X6=k1_relset_1(X6,X7,X8)|X6!=k1_xboole_0|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))&(X6!=k1_relset_1(X6,X7,X8)|v1_funct_2(X8,X6,X7)|X6!=k1_xboole_0|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))))&((~v1_funct_2(X8,X6,X7)|X8=k1_xboole_0|X6=k1_xboole_0|X7!=k1_xboole_0|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))&(X8!=k1_xboole_0|v1_funct_2(X8,X6,X7)|X6=k1_xboole_0|X7!=k1_xboole_0|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
fof(c_0_149, plain, ![X34, X35, X36]:(~v1_relat_1(X36)|(~r1_tarski(k1_relat_1(X36),X34)|~r1_tarski(k2_relat_1(X36),X35)|m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(X34,X35))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_relset_1])])).
cnf(c_0_150, negated_conjecture, (k1_relat_1(k2_relat_1(k2_relat_1(esk2_0)))=k1_relat_1(k2_relat_1(esk1_0))|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_142, c_0_115])).
cnf(c_0_151, negated_conjecture, (k2_relat_1(k1_relat_1(k1_relat_1(esk1_0)))=k1_relat_1(esk1_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~r1_tarski(esk1_0,k2_relat_1(esk1_0))|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_143, c_0_144])).
cnf(c_0_152, negated_conjecture, (k1_relat_1(k1_relat_1(esk1_0))=k2_relat_1(esk1_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~r1_tarski(esk1_0,k2_relat_1(esk1_0))|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_145, c_0_144])).
cnf(c_0_153, negated_conjecture, (~v1_funct_1(esk2_0)|~v1_funct_2(esk2_0,k1_relat_1(esk2_0),esk1_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(esk2_0),esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_154, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_155, negated_conjecture, (r1_tarski(k2_relat_1(esk2_0),X1)|~r1_tarski(esk1_0,k1_xboole_0)|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_146, c_0_147])).
cnf(c_0_156, plain, (v1_funct_2(X3,X1,X2)|X1!=k1_relset_1(X1,X2,X3)|X1!=k1_xboole_0|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_148])).
fof(c_0_157, plain, ![X26, X27, X28]:(~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(X26,X27)))|k1_relset_1(X26,X27,X28)=k1_relat_1(X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_158, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_159, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_149])).
cnf(c_0_160, negated_conjecture, (k2_relat_1(k1_relat_1(k2_relat_1(k2_relat_1(esk2_0))))=esk1_0|k1_relat_1(esk1_0)!=k1_xboole_0|~r1_tarski(esk1_0,k2_relat_1(k1_relat_1(k2_relat_1(k2_relat_1(esk2_0)))))|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_28, c_0_87])).
cnf(c_0_161, negated_conjecture, (k2_relat_1(k2_relat_1(k1_relat_1(k2_relat_1(esk1_0))))=k1_relat_1(k2_relat_1(esk1_0))|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_114, c_0_136])).
cnf(c_0_162, negated_conjecture, (k2_relat_1(k1_relat_1(k2_relat_1(esk1_0)))=k2_relat_1(k2_relat_1(esk2_0))|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_135, c_0_150])).
cnf(c_0_163, negated_conjecture, (k2_relat_1(k2_relat_1(esk1_0))=k1_relat_1(esk1_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~r1_tarski(esk1_0,k2_relat_1(esk1_0))|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_151, c_0_152])).
cnf(c_0_164, negated_conjecture, (~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(esk2_0),esk1_0)))|~v1_funct_2(esk2_0,k1_relat_1(esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_153, c_0_154])])).
cnf(c_0_165, negated_conjecture, (k2_relat_1(esk2_0)=k1_xboole_0|~r1_tarski(esk1_0,k1_xboole_0)|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_33, c_0_155])).
cnf(c_0_166, plain, (v1_funct_2(X1,k1_xboole_0,X2)|k1_relset_1(k1_xboole_0,X2,X1)!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,X2)))), inference(er,[status(thm)],[c_0_156])).
cnf(c_0_167, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_157])).
cnf(c_0_168, plain, (r1_tarski(X1,k2_zfmisc_1(X2,X3))|~r1_tarski(k2_relat_1(X1),X3)|~r1_tarski(k1_relat_1(X1),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_158, c_0_159])).
cnf(c_0_169, negated_conjecture, (k2_relat_1(k1_relat_1(k2_relat_1(esk1_0)))=esk1_0|k1_relat_1(esk1_0)!=k1_xboole_0|~r1_tarski(esk1_0,k2_relat_1(k1_relat_1(k2_relat_1(esk1_0))))|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_160, c_0_107])).
cnf(c_0_170, plain, (r1_tarski(X1,X2)|k1_relat_1(X3)!=k1_xboole_0|~r1_tarski(X1,k2_relat_1(X3))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_50, c_0_80])).
cnf(c_0_171, negated_conjecture, (k2_relat_1(k2_relat_1(k2_relat_1(esk2_0)))=k1_relat_1(k2_relat_1(esk1_0))|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_161, c_0_162])).
cnf(c_0_172, negated_conjecture, (k2_relat_1(k2_relat_1(esk2_0))=k1_relat_1(esk1_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~r1_tarski(esk1_0,k2_relat_1(esk1_0))|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_126, c_0_163])).
cnf(c_0_173, negated_conjecture, (~v1_funct_2(esk2_0,k1_relat_1(esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_159]), c_0_51]), c_0_69]), c_0_57])])).
cnf(c_0_174, negated_conjecture, (k1_relat_1(esk2_0)=k1_xboole_0|~r1_tarski(esk1_0,k1_xboole_0)|~v1_relat_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_165]), c_0_57])])).
cnf(c_0_175, plain, (v1_funct_2(X1,k1_xboole_0,X2)|k1_relat_1(X1)!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,X2)))), inference(spm,[status(thm)],[c_0_166, c_0_167])).
cnf(c_0_176, negated_conjecture, (r1_tarski(esk2_0,k2_zfmisc_1(X1,esk1_0))|~r1_tarski(k1_relat_1(esk2_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_51]), c_0_57])])).
cnf(c_0_177, negated_conjecture, (k2_relat_1(k1_relat_1(esk1_0))=esk1_0|k1_relat_1(esk1_0)!=k1_xboole_0|~r1_tarski(esk1_0,k2_relat_1(esk1_0))|~v1_relat_1(esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_144]), c_0_170])).
cnf(c_0_178, negated_conjecture, (k2_relat_1(k1_relat_1(esk1_0))=k1_relat_1(k2_relat_1(esk1_0))|k1_relat_1(esk1_0)!=k1_xboole_0|~r1_tarski(esk1_0,k2_relat_1(esk1_0))|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_171, c_0_172])).
cnf(c_0_179, negated_conjecture, (~v1_funct_2(esk2_0,k1_xboole_0,esk1_0)|~r1_tarski(esk1_0,k1_xboole_0)|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_173, c_0_174])).
cnf(c_0_180, plain, (v1_funct_2(X1,k1_xboole_0,X2)|~r1_tarski(k1_relat_1(X1),k1_xboole_0)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_159]), c_0_80]), c_0_33])).
cnf(c_0_181, negated_conjecture, (r1_tarski(esk2_0,k2_zfmisc_1(k1_relat_1(esk2_0),esk1_0))), inference(spm,[status(thm)],[c_0_176, c_0_69])).
cnf(c_0_182, plain, (k2_relat_1(X1)=k1_xboole_0|k1_relat_1(X2)!=k1_xboole_0|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_47]), c_0_29])])).
cnf(c_0_183, negated_conjecture, (k1_relat_1(k2_relat_1(esk1_0))=esk1_0|k1_relat_1(esk1_0)!=k1_xboole_0|~r1_tarski(esk1_0,k2_relat_1(esk1_0))|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_177, c_0_178])).
cnf(c_0_184, negated_conjecture, (~r1_tarski(k1_relat_1(esk2_0),k1_xboole_0)|~r1_tarski(esk1_0,k1_xboole_0)|~v1_relat_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_180]), c_0_57])])).
cnf(c_0_185, negated_conjecture, (r1_tarski(X1,k2_zfmisc_1(k1_relat_1(esk2_0),esk1_0))|~r1_tarski(X1,esk2_0)), inference(spm,[status(thm)],[c_0_50, c_0_181])).
cnf(c_0_186, plain, (v1_funct_2(X3,X1,X2)|X2=k1_xboole_0|X1!=k1_relset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_148])).
cnf(c_0_187, negated_conjecture, (k2_relat_1(k2_relat_1(esk2_0))=k1_xboole_0|k1_relat_1(esk1_0)!=k1_xboole_0|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_182, c_0_51])).
cnf(c_0_188, negated_conjecture, (k2_relat_1(k2_relat_1(esk1_0))=esk1_0|k1_relat_1(esk1_0)!=k1_xboole_0|~r1_tarski(esk1_0,k2_relat_1(esk1_0))|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_161, c_0_183])).
cnf(c_0_189, negated_conjecture, (~r1_tarski(esk1_0,k1_xboole_0)|~v1_relat_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_174]), c_0_29])])).
cnf(c_0_190, negated_conjecture, (v1_relat_1(X1)|~r1_tarski(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_185]), c_0_41])])).
cnf(c_0_191, plain, (X1=k1_xboole_0|v1_funct_2(X2,k1_relat_1(X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X2),X1)))), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_167])])).
cnf(c_0_192, negated_conjecture, (r1_tarski(k2_relat_1(X1),k1_xboole_0)|k1_relat_1(esk1_0)!=k1_xboole_0|~r1_tarski(X1,k2_relat_1(esk2_0))|~v1_relat_1(esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_187]), c_0_93])).
cnf(c_0_193, negated_conjecture, (k2_relat_1(k2_relat_1(esk2_0))=esk1_0|k1_relat_1(esk1_0)!=k1_xboole_0|~r1_tarski(esk1_0,k2_relat_1(esk1_0))|~v1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_126, c_0_188])).
cnf(c_0_194, negated_conjecture, (~r1_tarski(esk1_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_190]), c_0_63])).
cnf(c_0_195, plain, (X1=k1_xboole_0|v1_funct_2(X2,k1_relat_1(X2),X1)|~r1_tarski(X2,k2_zfmisc_1(k1_relat_1(X2),X1))), inference(spm,[status(thm)],[c_0_191, c_0_27])).
cnf(c_0_196, negated_conjecture, (k1_relat_1(esk1_0)!=k1_xboole_0|~r1_tarski(esk1_0,k2_relat_1(esk1_0))|~v1_relat_1(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_193]), c_0_69])]), c_0_194])).
cnf(c_0_197, negated_conjecture, (esk1_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_195]), c_0_181])])).
cnf(c_0_198, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_196, c_0_197]), c_0_73]), c_0_197]), c_0_197]), c_0_66]), c_0_29]), c_0_197]), c_0_67])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 199
# Proof object clause steps            : 167
# Proof object formula steps           : 32
# Proof object conjectures             : 93
# Proof object clause conjectures      : 90
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 135
# Proof object simplifying inferences  : 66
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 32
# Processed clauses                    : 12415
# ...of these trivial                  : 18
# ...subsumed                          : 10595
# ...remaining for further processing  : 1802
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 363
# Backward-rewritten                   : 756
# Generated clauses                    : 185864
# ...of the previous two non-trivial   : 151859
# Contextual simplify-reflections      : 205
# Paramodulations                      : 185857
# Factorizations                       : 0
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 646
#    Positive orientable unit clauses  : 26
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 619
# Current number of unprocessed clauses: 138223
# ...number of literals in the above   : 661606
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1150
# Clause-clause subsumption calls (NU) : 301642
# Rec. Clause-clause subsumption calls : 181595
# Non-unit clause-clause subsumptions  : 10508
# Unit Clause-clause subsumption calls : 1332
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 63
# BW rewrite match successes           : 8
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2602336

# -------------------------------------------------
# User time                : 1.431 s
# System time              : 0.020 s
# Total time               : 1.451 s
# Maximum resident set size: 95760 pages
