# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_mcart_1, axiom, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>![X2]:(X2=k2_mcart_1(X1)<=>![X3, X4]:(X1=k4_tarski(X3,X4)=>X2=X4))), file('data/MPTP/proofs/1bee52fe6f66473491af86f7810ecddc.E_input', d2_mcart_1)).
fof(t33_mcart_1, conjecture, ![X1, X2, X3, X4, X5, X6, X7, X8]:(k4_mcart_1(X1,X2,X3,X4)=k4_mcart_1(X5,X6,X7,X8)=>(X1=X5&(X2=X6&(X3=X7&X4=X8)))), file('data/MPTP/proofs/1bee52fe6f66473491af86f7810ecddc.E_input', t33_mcart_1)).
fof(t31_mcart_1, axiom, ![X1, X2, X3, X4]:k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4), file('data/MPTP/proofs/1bee52fe6f66473491af86f7810ecddc.E_input', t31_mcart_1)).
fof(d1_mcart_1, axiom, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>![X2]:(X2=k1_mcart_1(X1)<=>![X3, X4]:(X1=k4_tarski(X3,X4)=>X2=X3))), file('data/MPTP/proofs/1bee52fe6f66473491af86f7810ecddc.E_input', d1_mcart_1)).
fof(c_0_4, plain, ![X9, X10, X11, X12, X13, X14, X15]:((X12!=k2_mcart_1(X9)|(X9!=k4_tarski(X13,X14)|X12=X14)|X9!=k4_tarski(X10,X11))&((X9=k4_tarski(esk1_2(X9,X15),esk2_2(X9,X15))|X15=k2_mcart_1(X9)|X9!=k4_tarski(X10,X11))&(X15!=esk2_2(X9,X15)|X15=k2_mcart_1(X9)|X9!=k4_tarski(X10,X11)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_mcart_1])])])])])])).
fof(c_0_5, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6, X7, X8]:(k4_mcart_1(X1,X2,X3,X4)=k4_mcart_1(X5,X6,X7,X8)=>(X1=X5&(X2=X6&(X3=X7&X4=X8))))), inference(assume_negation,[status(cth)],[t33_mcart_1])).
cnf(c_0_6, plain, (X1=X4|X1!=k2_mcart_1(X2)|X2!=k4_tarski(X3,X4)|X2!=k4_tarski(X5,X6)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_7, negated_conjecture, (k4_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0)=k4_mcart_1(esk7_0,esk8_0,esk9_0,esk10_0)&(esk3_0!=esk7_0|(esk4_0!=esk8_0|(esk5_0!=esk9_0|esk6_0!=esk10_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
fof(c_0_8, plain, ![X26, X27, X28, X29]:k4_mcart_1(X26,X27,X28,X29)=k4_tarski(k4_tarski(k4_tarski(X26,X27),X28),X29), inference(variable_rename,[status(thm)],[t31_mcart_1])).
fof(c_0_9, plain, ![X30, X31, X32, X33, X34, X35, X36]:((X33!=k1_mcart_1(X30)|(X30!=k4_tarski(X34,X35)|X33=X34)|X30!=k4_tarski(X31,X32))&((X30=k4_tarski(esk11_2(X30,X36),esk12_2(X30,X36))|X36=k1_mcart_1(X30)|X30!=k4_tarski(X31,X32))&(X36!=esk11_2(X30,X36)|X36=k1_mcart_1(X30)|X30!=k4_tarski(X31,X32)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_mcart_1])])])])])])).
cnf(c_0_10, plain, (k2_mcart_1(k4_tarski(X1,X2))=X3|k4_tarski(X1,X2)!=k4_tarski(X4,X3)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_6])])).
cnf(c_0_11, negated_conjecture, (k4_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0)=k4_mcart_1(esk7_0,esk8_0,esk9_0,esk10_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (X1=X3|X1!=k1_mcart_1(X2)|X2!=k4_tarski(X3,X4)|X2!=k4_tarski(X5,X6)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(er,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (k4_tarski(k4_tarski(k4_tarski(esk7_0,esk8_0),esk9_0),esk10_0)=k4_tarski(k4_tarski(k4_tarski(esk3_0,esk4_0),esk5_0),esk6_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_12]), c_0_12])).
cnf(c_0_16, plain, (k1_mcart_1(k4_tarski(X1,X2))=X3|k4_tarski(X1,X2)!=k4_tarski(X3,X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_13])])).
cnf(c_0_17, negated_conjecture, (esk6_0=esk10_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_14])).
cnf(c_0_18, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_19, negated_conjecture, (k4_tarski(k4_tarski(k4_tarski(esk7_0,esk8_0),esk9_0),esk10_0)=k4_tarski(k4_tarski(k4_tarski(esk3_0,esk4_0),esk5_0),esk10_0)), inference(rw,[status(thm)],[c_0_15, c_0_17])).
cnf(c_0_20, negated_conjecture, (k4_tarski(k4_tarski(esk7_0,esk8_0),esk9_0)=k4_tarski(k4_tarski(esk3_0,esk4_0),esk5_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_18])).
cnf(c_0_21, negated_conjecture, (esk5_0=esk9_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_20]), c_0_14])).
cnf(c_0_22, negated_conjecture, (k4_tarski(k4_tarski(esk7_0,esk8_0),esk9_0)=k4_tarski(k4_tarski(esk3_0,esk4_0),esk9_0)), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_23, negated_conjecture, (esk3_0!=esk7_0|esk4_0!=esk8_0|esk5_0!=esk9_0|esk6_0!=esk10_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_24, negated_conjecture, (k4_tarski(esk7_0,esk8_0)=k4_tarski(esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_22]), c_0_18])).
cnf(c_0_25, negated_conjecture, (esk7_0!=esk3_0|esk4_0!=esk8_0|esk5_0!=esk9_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_17])])).
cnf(c_0_26, negated_conjecture, (esk4_0=esk8_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_24]), c_0_14])).
cnf(c_0_27, negated_conjecture, (esk7_0!=esk3_0|esk4_0!=esk8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_21])])).
cnf(c_0_28, negated_conjecture, (k4_tarski(esk7_0,esk8_0)=k4_tarski(esk3_0,esk8_0)), inference(rw,[status(thm)],[c_0_24, c_0_26])).
cnf(c_0_29, negated_conjecture, (esk7_0!=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_26])])).
cnf(c_0_30, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_28]), c_0_18]), c_0_29]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 31
# Proof object clause steps            : 22
# Proof object formula steps           : 9
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 5
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 9
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 8
# Processed clauses                    : 42
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 40
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 14
# Generated clauses                    : 97
# ...of the previous two non-trivial   : 97
# Contextual simplify-reflections      : 0
# Paramodulations                      : 87
# Factorizations                       : 0
# Equation resolutions                 : 12
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 12
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 5
# Current number of unprocessed clauses: 60
# ...number of literals in the above   : 146
# Current number of archived formulas  : 0
# Current number of archived clauses   : 23
# Clause-clause subsumption calls (NU) : 33
# Rec. Clause-clause subsumption calls : 31
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1528

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.000 s
# Total time               : 0.009 s
# Maximum resident set size: 28852 pages
