# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/1c0aeac6a3894644acffcdc1f8b1683d.E_input', t4_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/1c0aeac6a3894644acffcdc1f8b1683d.E_input', t28_xboole_1)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/1c0aeac6a3894644acffcdc1f8b1683d.E_input', l54_zfmisc_1)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/1c0aeac6a3894644acffcdc1f8b1683d.E_input', t65_xboole_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/1c0aeac6a3894644acffcdc1f8b1683d.E_input', t113_zfmisc_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/1c0aeac6a3894644acffcdc1f8b1683d.E_input', t2_xboole_1)).
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/1c0aeac6a3894644acffcdc1f8b1683d.E_input', d2_zfmisc_1)).
fof(t138_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=>(k2_zfmisc_1(X1,X2)=k1_xboole_0|(r1_tarski(X1,X3)&r1_tarski(X2,X4)))), file('data/MPTP/proofs/1c0aeac6a3894644acffcdc1f8b1683d.E_input', t138_zfmisc_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/1c0aeac6a3894644acffcdc1f8b1683d.E_input', d3_tarski)).
fof(c_0_9, plain, ![X7, X8, X10, X11, X12]:((r1_xboole_0(X7,X8)|r2_hidden(esk1_2(X7,X8),k3_xboole_0(X7,X8)))&(~r2_hidden(X12,k3_xboole_0(X10,X11))|~r1_xboole_0(X10,X11))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_10, plain, ![X13, X14]:(~r1_tarski(X13,X14)|k3_xboole_0(X13,X14)=X13), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_11, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_13, plain, ![X15, X16, X17, X18]:(((r2_hidden(X15,X17)|~r2_hidden(k4_tarski(X15,X16),k2_zfmisc_1(X17,X18)))&(r2_hidden(X16,X18)|~r2_hidden(k4_tarski(X15,X16),k2_zfmisc_1(X17,X18))))&(~r2_hidden(X15,X17)|~r2_hidden(X16,X18)|r2_hidden(k4_tarski(X15,X16),k2_zfmisc_1(X17,X18)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
cnf(c_0_14, plain, (~r1_tarski(X1,X2)|~r2_hidden(X3,X1)|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_15, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_16, plain, ![X19]:r1_xboole_0(X19,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
fof(c_0_17, plain, ![X20, X21]:((k2_zfmisc_1(X20,X21)!=k1_xboole_0|(X20=k1_xboole_0|X21=k1_xboole_0))&((X20!=k1_xboole_0|k2_zfmisc_1(X20,X21)=k1_xboole_0)&(X21!=k1_xboole_0|k2_zfmisc_1(X20,X21)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_18, plain, (~r1_tarski(k2_zfmisc_1(X1,X2),X3)|~r2_hidden(X4,X2)|~r2_hidden(X5,X1)|~r1_xboole_0(k2_zfmisc_1(X1,X2),X3)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_19, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_21, plain, ![X22]:r1_tarski(k1_xboole_0,X22), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_22, plain, (~r1_tarski(k2_zfmisc_1(X1,X2),k1_xboole_0)|~r2_hidden(X3,X2)|~r2_hidden(X4,X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_23, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_24, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_25, plain, ![X23, X24, X25, X26, X29, X30, X31, X32, X33, X34, X36, X37]:((((r2_hidden(esk2_4(X23,X24,X25,X26),X23)|~r2_hidden(X26,X25)|X25!=k2_zfmisc_1(X23,X24))&((r2_hidden(esk3_4(X23,X24,X25,X26),X24)|~r2_hidden(X26,X25)|X25!=k2_zfmisc_1(X23,X24))&(X26=k4_tarski(esk2_4(X23,X24,X25,X26),esk3_4(X23,X24,X25,X26))|~r2_hidden(X26,X25)|X25!=k2_zfmisc_1(X23,X24))))&(~r2_hidden(X30,X23)|(~r2_hidden(X31,X24)|X29!=k4_tarski(X30,X31))|r2_hidden(X29,X25)|X25!=k2_zfmisc_1(X23,X24)))&((~r2_hidden(esk4_3(X32,X33,X34),X34)|(~r2_hidden(X36,X32)|(~r2_hidden(X37,X33)|esk4_3(X32,X33,X34)!=k4_tarski(X36,X37)))|X34=k2_zfmisc_1(X32,X33))&((r2_hidden(esk5_3(X32,X33,X34),X32)|r2_hidden(esk4_3(X32,X33,X34),X34)|X34=k2_zfmisc_1(X32,X33))&((r2_hidden(esk6_3(X32,X33,X34),X33)|r2_hidden(esk4_3(X32,X33,X34),X34)|X34=k2_zfmisc_1(X32,X33))&(esk4_3(X32,X33,X34)=k4_tarski(esk5_3(X32,X33,X34),esk6_3(X32,X33,X34))|r2_hidden(esk4_3(X32,X33,X34),X34)|X34=k2_zfmisc_1(X32,X33)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
cnf(c_0_26, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24])])).
cnf(c_0_27, plain, (r2_hidden(esk2_4(X1,X2,X3,X4),X1)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_28, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X3)|~r2_hidden(X4,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_23])).
cnf(c_0_29, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk1_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_30, negated_conjecture, ~(![X1, X2, X3, X4]:(r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=>(k2_zfmisc_1(X1,X2)=k1_xboole_0|(r1_tarski(X1,X3)&r1_tarski(X2,X4))))), inference(assume_negation,[status(cth)],[t138_zfmisc_1])).
cnf(c_0_31, plain, (X1!=k2_zfmisc_1(X2,X3)|~r1_tarski(X2,X4)|~r2_hidden(X5,X1)|~r1_xboole_0(X2,X4)), inference(spm,[status(thm)],[c_0_14, c_0_27])).
cnf(c_0_32, plain, (r1_xboole_0(X1,X2)|X3!=k1_xboole_0|~r2_hidden(X4,X3)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
fof(c_0_33, plain, ![X44, X45, X46, X47, X48]:((~r1_tarski(X44,X45)|(~r2_hidden(X46,X44)|r2_hidden(X46,X45)))&((r2_hidden(esk11_2(X47,X48),X47)|r1_tarski(X47,X48))&(~r2_hidden(esk11_2(X47,X48),X48)|r1_tarski(X47,X48)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_34, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk7_0,esk8_0),k2_zfmisc_1(esk9_0,esk10_0))&(k2_zfmisc_1(esk7_0,esk8_0)!=k1_xboole_0&(~r1_tarski(esk7_0,esk9_0)|~r1_tarski(esk8_0,esk10_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])).
cnf(c_0_35, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_23]), c_0_24])]), c_0_32])).
cnf(c_0_36, plain, (r2_hidden(esk3_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_37, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_38, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk7_0,esk8_0),k2_zfmisc_1(esk9_0,esk10_0))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_39, plain, (X1!=k2_zfmisc_1(X2,X3)|X3!=k1_xboole_0|~r2_hidden(X4,X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_40, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_41, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,X2))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_42, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(esk9_0,esk10_0))|~r2_hidden(X1,k2_zfmisc_1(esk7_0,esk8_0))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_43, plain, (X1!=k1_xboole_0|~r2_hidden(X2,k2_zfmisc_1(X3,X1))), inference(er,[status(thm)],[c_0_39])).
cnf(c_0_44, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_45, negated_conjecture, (r2_hidden(X1,esk10_0)|~r2_hidden(k4_tarski(X2,X1),k2_zfmisc_1(esk7_0,esk8_0))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_46, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_47, plain, (r2_hidden(esk5_3(X1,X2,X3),X1)|r2_hidden(esk4_3(X1,X2,X3),X3)|X3=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_48, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_49, negated_conjecture, (r2_hidden(X1,esk10_0)|~r2_hidden(X1,esk8_0)|~r2_hidden(X2,esk7_0)), inference(spm,[status(thm)],[c_0_45, c_0_15])).
cnf(c_0_50, plain, (X1=k1_xboole_0|r2_hidden(esk4_3(k1_xboole_0,X2,X1),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_23])).
cnf(c_0_51, negated_conjecture, (r2_hidden(X1,esk9_0)|~r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(esk7_0,esk8_0))), inference(spm,[status(thm)],[c_0_48, c_0_42])).
cnf(c_0_52, plain, (r1_tarski(X1,X2)|~r2_hidden(esk11_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_53, negated_conjecture, (esk7_0=k1_xboole_0|r2_hidden(X1,esk10_0)|~r2_hidden(X1,esk8_0)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_54, negated_conjecture, (r2_hidden(X1,esk9_0)|~r2_hidden(X2,esk8_0)|~r2_hidden(X1,esk7_0)), inference(spm,[status(thm)],[c_0_51, c_0_15])).
cnf(c_0_55, negated_conjecture, (esk7_0=k1_xboole_0|r1_tarski(X1,esk10_0)|~r2_hidden(esk11_2(X1,esk10_0),esk8_0)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_56, plain, (r2_hidden(esk11_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_57, negated_conjecture, (esk8_0=k1_xboole_0|r2_hidden(X1,esk9_0)|~r2_hidden(X1,esk7_0)), inference(spm,[status(thm)],[c_0_54, c_0_50])).
cnf(c_0_58, negated_conjecture, (~r1_tarski(esk7_0,esk9_0)|~r1_tarski(esk8_0,esk10_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_59, negated_conjecture, (esk7_0=k1_xboole_0|r1_tarski(esk8_0,esk10_0)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_60, negated_conjecture, (esk8_0=k1_xboole_0|r1_tarski(X1,esk9_0)|~r2_hidden(esk11_2(X1,esk9_0),esk7_0)), inference(spm,[status(thm)],[c_0_52, c_0_57])).
cnf(c_0_61, negated_conjecture, (esk7_0=k1_xboole_0|~r1_tarski(esk7_0,esk9_0)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_62, negated_conjecture, (esk8_0=k1_xboole_0|r1_tarski(esk7_0,esk9_0)), inference(spm,[status(thm)],[c_0_60, c_0_56])).
cnf(c_0_63, negated_conjecture, (k2_zfmisc_1(esk7_0,esk8_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_64, negated_conjecture, (esk8_0=k1_xboole_0|esk7_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_65, negated_conjecture, (esk7_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_44])])).
cnf(c_0_66, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_65]), c_0_23])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 67
# Proof object clause steps            : 48
# Proof object formula steps           : 19
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 28
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 25
# Processed clauses                    : 217
# ...of these trivial                  : 1
# ...subsumed                          : 68
# ...remaining for further processing  : 148
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 18
# Backward-rewritten                   : 35
# Generated clauses                    : 327
# ...of the previous two non-trivial   : 315
# Contextual simplify-reflections      : 3
# Paramodulations                      : 317
# Factorizations                       : 0
# Equation resolutions                 : 10
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 70
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 8
#    Non-unit-clauses                  : 55
# Current number of unprocessed clauses: 125
# ...number of literals in the above   : 468
# Current number of archived formulas  : 0
# Current number of archived clauses   : 78
# Clause-clause subsumption calls (NU) : 785
# Rec. Clause-clause subsumption calls : 484
# Non-unit clause-clause subsumptions  : 70
# Unit Clause-clause subsumption calls : 145
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5228

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.006 s
# Total time               : 0.013 s
# Maximum resident set size: 28776 pages
