# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/1c175dfa4e64482b9e7ac46429c8b415.E_input', t28_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/1c175dfa4e64482b9e7ac46429c8b415.E_input', t48_xboole_1)).
fof(t148_zfmisc_1, conjecture, ![X1, X2, X3, X4]:((r1_tarski(X1,X2)&(k3_xboole_0(X2,X3)=k1_tarski(X4)&r2_hidden(X4,X1)))=>k3_xboole_0(X1,X3)=k1_tarski(X4)), file('data/MPTP/proofs/1c175dfa4e64482b9e7ac46429c8b415.E_input', t148_zfmisc_1)).
fof(l33_zfmisc_1, axiom, ![X1, X2]:(k4_xboole_0(k1_tarski(X1),X2)=k1_tarski(X1)<=>~(r2_hidden(X1,X2))), file('data/MPTP/proofs/1c175dfa4e64482b9e7ac46429c8b415.E_input', l33_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/1c175dfa4e64482b9e7ac46429c8b415.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/1c175dfa4e64482b9e7ac46429c8b415.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/1c175dfa4e64482b9e7ac46429c8b415.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/1c175dfa4e64482b9e7ac46429c8b415.E_input', t72_enumset1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/1c175dfa4e64482b9e7ac46429c8b415.E_input', d5_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/1c175dfa4e64482b9e7ac46429c8b415.E_input', commutativity_k3_xboole_0)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/1c175dfa4e64482b9e7ac46429c8b415.E_input', t3_boole)).
fof(t47_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/1c175dfa4e64482b9e7ac46429c8b415.E_input', t47_xboole_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/1c175dfa4e64482b9e7ac46429c8b415.E_input', d3_tarski)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/1c175dfa4e64482b9e7ac46429c8b415.E_input', d10_xboole_0)).
fof(t49_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),X3), file('data/MPTP/proofs/1c175dfa4e64482b9e7ac46429c8b415.E_input', t49_xboole_1)).
fof(c_0_15, plain, ![X5, X6]:(~r1_tarski(X5,X6)|k3_xboole_0(X5,X6)=X5), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_16, plain, ![X7, X8]:k4_xboole_0(X7,k4_xboole_0(X7,X8))=k3_xboole_0(X7,X8), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_17, negated_conjecture, ~(![X1, X2, X3, X4]:((r1_tarski(X1,X2)&(k3_xboole_0(X2,X3)=k1_tarski(X4)&r2_hidden(X4,X1)))=>k3_xboole_0(X1,X3)=k1_tarski(X4))), inference(assume_negation,[status(cth)],[t148_zfmisc_1])).
fof(c_0_18, plain, ![X13, X14]:((k4_xboole_0(k1_tarski(X13),X14)!=k1_tarski(X13)|~r2_hidden(X13,X14))&(r2_hidden(X13,X14)|k4_xboole_0(k1_tarski(X13),X14)=k1_tarski(X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l33_zfmisc_1])])])).
fof(c_0_19, plain, ![X15]:k2_tarski(X15,X15)=k1_tarski(X15), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_20, plain, ![X16, X17]:k1_enumset1(X16,X16,X17)=k2_tarski(X16,X17), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_21, plain, ![X18, X19, X20]:k2_enumset1(X18,X18,X19,X20)=k1_enumset1(X18,X19,X20), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_22, plain, ![X21, X22, X23, X24]:k3_enumset1(X21,X21,X22,X23,X24)=k2_enumset1(X21,X22,X23,X24), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_23, plain, ![X25, X26, X27, X28, X29, X30, X31, X32]:((((r2_hidden(X28,X25)|~r2_hidden(X28,X27)|X27!=k4_xboole_0(X25,X26))&(~r2_hidden(X28,X26)|~r2_hidden(X28,X27)|X27!=k4_xboole_0(X25,X26)))&(~r2_hidden(X29,X25)|r2_hidden(X29,X26)|r2_hidden(X29,X27)|X27!=k4_xboole_0(X25,X26)))&((~r2_hidden(esk5_3(X30,X31,X32),X32)|(~r2_hidden(esk5_3(X30,X31,X32),X30)|r2_hidden(esk5_3(X30,X31,X32),X31))|X32=k4_xboole_0(X30,X31))&((r2_hidden(esk5_3(X30,X31,X32),X30)|r2_hidden(esk5_3(X30,X31,X32),X32)|X32=k4_xboole_0(X30,X31))&(~r2_hidden(esk5_3(X30,X31,X32),X31)|r2_hidden(esk5_3(X30,X31,X32),X32)|X32=k4_xboole_0(X30,X31))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_24, plain, ![X34, X35]:k3_xboole_0(X34,X35)=k3_xboole_0(X35,X34), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_25, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_26, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_27, negated_conjecture, ((r1_tarski(esk1_0,esk2_0)&(k3_xboole_0(esk2_0,esk3_0)=k1_tarski(esk4_0)&r2_hidden(esk4_0,esk1_0)))&k3_xboole_0(esk1_0,esk3_0)!=k1_tarski(esk4_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).
cnf(c_0_28, plain, (k4_xboole_0(k1_tarski(X1),X2)!=k1_tarski(X1)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_29, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_30, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_31, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_32, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_33, plain, ![X36]:k4_xboole_0(X36,k1_xboole_0)=X36, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_34, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_35, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_36, plain, ![X37, X38]:k4_xboole_0(X37,k3_xboole_0(X37,X38))=k4_xboole_0(X37,X38), inference(variable_rename,[status(thm)],[t47_xboole_1])).
cnf(c_0_37, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=X1|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_38, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_39, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_40, plain, (k4_xboole_0(k3_enumset1(X1,X1,X1,X1,X1),X2)!=k3_enumset1(X1,X1,X1,X1,X1)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_29]), c_0_29]), c_0_30]), c_0_30]), c_0_31]), c_0_31]), c_0_32]), c_0_32])).
cnf(c_0_41, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_42, plain, ![X39, X40, X41, X42, X43]:((~r1_tarski(X39,X40)|(~r2_hidden(X41,X39)|r2_hidden(X41,X40)))&((r2_hidden(esk6_2(X42,X43),X42)|r1_tarski(X42,X43))&(~r2_hidden(esk6_2(X42,X43),X43)|r1_tarski(X42,X43)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_43, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_34])).
cnf(c_0_44, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_26]), c_0_26])).
cnf(c_0_45, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_46, negated_conjecture, (k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0))=esk1_0), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_47, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_39])).
fof(c_0_48, plain, ![X45, X46]:(((r1_tarski(X45,X46)|X45!=X46)&(r1_tarski(X46,X45)|X45!=X46))&(~r1_tarski(X45,X46)|~r1_tarski(X46,X45)|X45=X46)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_49, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_50, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_51, plain, (~r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))|~r2_hidden(X1,k4_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_52, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X1,X2)))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_45, c_0_26])).
cnf(c_0_53, negated_conjecture, (~r2_hidden(X1,k4_xboole_0(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_46]), c_0_47])).
fof(c_0_54, plain, ![X47, X48, X49]:k3_xboole_0(X47,k4_xboole_0(X48,X49))=k4_xboole_0(k3_xboole_0(X47,X48),X49), inference(variable_rename,[status(thm)],[t49_xboole_1])).
cnf(c_0_55, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_56, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_57, plain, (~r2_hidden(X1,k4_xboole_0(k4_xboole_0(X2,X3),X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_47])).
cnf(c_0_58, negated_conjecture, (r1_tarski(k4_xboole_0(esk1_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_53, c_0_50])).
cnf(c_0_59, negated_conjecture, (k4_xboole_0(esk1_0,esk2_0)=k4_xboole_0(esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_52, c_0_46])).
cnf(c_0_60, negated_conjecture, (k3_xboole_0(esk2_0,esk3_0)=k1_tarski(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_61, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_62, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_63, plain, (r1_tarski(k4_xboole_0(k4_xboole_0(X1,X2),X1),X3)), inference(spm,[status(thm)],[c_0_57, c_0_50])).
cnf(c_0_64, negated_conjecture, (r1_tarski(k4_xboole_0(esk1_0,esk1_0),X1)), inference(rw,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_65, negated_conjecture, (k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0))=k3_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_29]), c_0_30]), c_0_31]), c_0_26]), c_0_32])).
cnf(c_0_66, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,X3)))=k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_26]), c_0_26])).
cnf(c_0_67, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_68, negated_conjecture, (k4_xboole_0(esk1_0,esk1_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_62, c_0_64])).
cnf(c_0_69, negated_conjecture, (k4_xboole_0(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0)),X1)!=k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0))|~r2_hidden(esk4_0,X1)), inference(spm,[status(thm)],[c_0_40, c_0_65])).
cnf(c_0_70, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,X1)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_71, negated_conjecture, (k4_xboole_0(esk1_0,esk2_0)=k1_xboole_0), inference(rw,[status(thm)],[c_0_59, c_0_68])).
cnf(c_0_72, negated_conjecture, (k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,k4_xboole_0(esk3_0,X1)))!=k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0))|~r2_hidden(esk4_0,X1)), inference(rw,[status(thm)],[c_0_69, c_0_66])).
cnf(c_0_73, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_70]), c_0_41])).
cnf(c_0_74, negated_conjecture, (k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,k4_xboole_0(esk2_0,X1)))=k4_xboole_0(esk1_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_71]), c_0_41])).
cnf(c_0_75, negated_conjecture, (~r2_hidden(esk4_0,k4_xboole_0(X1,esk3_0))), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_76, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_77, negated_conjecture, (k4_xboole_0(esk1_0,k4_xboole_0(esk2_0,X1))=k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,X1))), inference(spm,[status(thm)],[c_0_52, c_0_74])).
cnf(c_0_78, plain, (~r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,k4_xboole_0(X3,X4))))|~r2_hidden(X1,X4)), inference(spm,[status(thm)],[c_0_43, c_0_66])).
cnf(c_0_79, negated_conjecture, (~r2_hidden(esk4_0,k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,esk3_0))))), inference(spm,[status(thm)],[c_0_75, c_0_66])).
cnf(c_0_80, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_76])).
cnf(c_0_81, negated_conjecture, (k4_xboole_0(esk1_0,k4_xboole_0(X1,k4_xboole_0(X1,esk2_0)))=k4_xboole_0(esk1_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_44]), c_0_77]), c_0_52])).
cnf(c_0_82, plain, (~r2_hidden(X1,k4_xboole_0(k4_xboole_0(X2,X3),X4))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_67]), c_0_41])).
cnf(c_0_83, negated_conjecture, (~r2_hidden(esk4_0,k4_xboole_0(k4_xboole_0(X1,esk3_0),X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_67]), c_0_41])).
cnf(c_0_84, negated_conjecture, (r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,esk2_0)))|r2_hidden(X1,k4_xboole_0(esk1_0,X2))|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_85, negated_conjecture, (r2_hidden(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_86, plain, (r2_hidden(X1,X2)|k4_xboole_0(k1_tarski(X1),X2)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_87, negated_conjecture, (k3_xboole_0(esk1_0,esk3_0)!=k1_tarski(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_88, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,k4_xboole_0(X4,X2))), inference(spm,[status(thm)],[c_0_82, c_0_73])).
cnf(c_0_89, negated_conjecture, (r2_hidden(esk4_0,k4_xboole_0(esk1_0,k4_xboole_0(X1,esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_85])])).
cnf(c_0_90, plain, (k4_xboole_0(k3_enumset1(X1,X1,X1,X1,X1),X2)=k3_enumset1(X1,X1,X1,X1,X1)|r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_29]), c_0_29]), c_0_30]), c_0_30]), c_0_31]), c_0_31]), c_0_32]), c_0_32])).
cnf(c_0_91, negated_conjecture, (k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))!=k3_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_29]), c_0_30]), c_0_31]), c_0_26]), c_0_32])).
cnf(c_0_92, negated_conjecture, (~r2_hidden(esk4_0,k4_xboole_0(X1,esk1_0))), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_93, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k3_enumset1(X2,X2,X2,X2,X2)))=k3_enumset1(X2,X2,X2,X2,X2)|r2_hidden(X2,k4_xboole_0(k3_enumset1(X2,X2,X2,X2,X2),X1))), inference(spm,[status(thm)],[c_0_44, c_0_90])).
cnf(c_0_94, negated_conjecture, (k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0))!=k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))), inference(rw,[status(thm)],[c_0_91, c_0_65])).
cnf(c_0_95, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_65]), c_0_77]), c_0_77]), c_0_52]), c_0_65]), c_0_94]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 96
# Proof object clause steps            : 65
# Proof object formula steps           : 31
# Proof object conjectures             : 29
# Proof object clause conjectures      : 26
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 29
# Proof object simplifying inferences  : 55
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 23
# Processed clauses                    : 1339
# ...of these trivial                  : 42
# ...subsumed                          : 1032
# ...remaining for further processing  : 265
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 9
# Backward-rewritten                   : 52
# Generated clauses                    : 8352
# ...of the previous two non-trivial   : 6359
# Contextual simplify-reflections      : 3
# Paramodulations                      : 8341
# Factorizations                       : 6
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 177
#    Positive orientable unit clauses  : 56
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 22
#    Non-unit-clauses                  : 96
# Current number of unprocessed clauses: 4850
# ...number of literals in the above   : 10653
# Current number of archived formulas  : 0
# Current number of archived clauses   : 88
# Clause-clause subsumption calls (NU) : 2238
# Rec. Clause-clause subsumption calls : 1843
# Non-unit clause-clause subsumptions  : 294
# Unit Clause-clause subsumption calls : 270
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 546
# BW rewrite match successes           : 112
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 116862

# -------------------------------------------------
# User time                : 0.066 s
# System time              : 0.003 s
# Total time               : 0.070 s
# Maximum resident set size: 28592 pages
