# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.011 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d13_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))=>((r1_tsep_1(X3,X4)|r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6)))=>![X7]:((v1_funct_1(X7)&(v1_funct_2(X7,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))))=>(X7=k10_tmap_1(X1,X2,X3,X4,X5,X6)<=>(r2_funct_2(u1_struct_0(X3),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X7),X5)&r2_funct_2(u1_struct_0(X4),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X7),X6))))))))))), file('data/MPTP/proofs/1c21b66f7934468a83816078da9ff39d.E_input', d13_tmap_1)).
fof(dt_k10_tmap_1, axiom, ![X1, X2, X3, X4, X5, X6]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))&((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))&((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))&((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))&(v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))))))=>(v1_funct_1(k10_tmap_1(X1,X2,X3,X4,X5,X6))&(v1_funct_2(k10_tmap_1(X1,X2,X3,X4,X5,X6),u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&m1_subset_1(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))))), file('data/MPTP/proofs/1c21b66f7934468a83816078da9ff39d.E_input', dt_k10_tmap_1)).
fof(t144_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(~(r1_tsep_1(X3,X4))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))=>((r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6))&r4_tsep_1(X1,X3,X4))=>(v1_funct_1(k10_tmap_1(X1,X2,X3,X4,X5,X6))&(v1_funct_2(k10_tmap_1(X1,X2,X3,X4,X5,X6),u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&(v5_pre_topc(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_tsep_1(X1,X3,X4),X2)&m1_subset_1(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))))))))))))), file('data/MPTP/proofs/1c21b66f7934468a83816078da9ff39d.E_input', t144_tmap_1)).
fof(redefinition_r2_funct_2, axiom, ![X1, X2, X3, X4]:(((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&(v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))=>(r2_funct_2(X1,X2,X3,X4)<=>X3=X4)), file('data/MPTP/proofs/1c21b66f7934468a83816078da9ff39d.E_input', redefinition_r2_funct_2)).
fof(dt_k3_tmap_1, axiom, ![X1, X2, X3, X4, X5]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))&(m1_pre_topc(X3,X1)&(m1_pre_topc(X4,X1)&(v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))))))=>(v1_funct_1(k3_tmap_1(X1,X2,X3,X4,X5))&(v1_funct_2(k3_tmap_1(X1,X2,X3,X4,X5),u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(k3_tmap_1(X1,X2,X3,X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))), file('data/MPTP/proofs/1c21b66f7934468a83816078da9ff39d.E_input', dt_k3_tmap_1)).
fof(t126_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(r4_tsep_1(X1,X3,X4)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))))=>((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&(v5_pre_topc(X5,k1_tsep_1(X1,X3,X4),X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2)))))))<=>((v1_funct_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5))&(v1_funct_2(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),X3,X2)&m1_subset_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))&(v1_funct_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5))&(v1_funct_2(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5),u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5),X4,X2)&m1_subset_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))))))))))), file('data/MPTP/proofs/1c21b66f7934468a83816078da9ff39d.E_input', t126_tmap_1)).
fof(t2_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>m1_pre_topc(X1,X1)), file('data/MPTP/proofs/1c21b66f7934468a83816078da9ff39d.E_input', t2_tsep_1)).
fof(t29_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>((m1_pre_topc(X2,X3)&m1_pre_topc(X4,X3))=>m1_pre_topc(k1_tsep_1(X1,X2,X4),X3)))))), file('data/MPTP/proofs/1c21b66f7934468a83816078da9ff39d.E_input', t29_tmap_1)).
fof(c_0_8, plain, ![X19, X20, X21, X22, X23, X24, X25]:((((r2_funct_2(u1_struct_0(X21),u1_struct_0(X20),k3_tmap_1(X19,X20,k1_tsep_1(X19,X21,X22),X21,X25),X23)|X25!=k10_tmap_1(X19,X20,X21,X22,X23,X24)|(~v1_funct_1(X25)|(~v1_funct_2(X25,u1_struct_0(k1_tsep_1(X19,X21,X22)),u1_struct_0(X20))|~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X19,X21,X22)),u1_struct_0(X20))))))|~r1_tsep_1(X21,X22)|(~v1_funct_1(X24)|(~v1_funct_2(X24,u1_struct_0(X22),u1_struct_0(X20))|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(X20))))))|(~v1_funct_1(X23)|(~v1_funct_2(X23,u1_struct_0(X21),u1_struct_0(X20))|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X21),u1_struct_0(X20))))))|(v2_struct_0(X22)|~m1_pre_topc(X22,X19))|(v2_struct_0(X21)|~m1_pre_topc(X21,X19))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20)))|(v2_struct_0(X19)|(~v2_pre_topc(X19)|~l1_pre_topc(X19))))&(r2_funct_2(u1_struct_0(X22),u1_struct_0(X20),k3_tmap_1(X19,X20,k1_tsep_1(X19,X21,X22),X22,X25),X24)|X25!=k10_tmap_1(X19,X20,X21,X22,X23,X24)|(~v1_funct_1(X25)|(~v1_funct_2(X25,u1_struct_0(k1_tsep_1(X19,X21,X22)),u1_struct_0(X20))|~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X19,X21,X22)),u1_struct_0(X20))))))|~r1_tsep_1(X21,X22)|(~v1_funct_1(X24)|(~v1_funct_2(X24,u1_struct_0(X22),u1_struct_0(X20))|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(X20))))))|(~v1_funct_1(X23)|(~v1_funct_2(X23,u1_struct_0(X21),u1_struct_0(X20))|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X21),u1_struct_0(X20))))))|(v2_struct_0(X22)|~m1_pre_topc(X22,X19))|(v2_struct_0(X21)|~m1_pre_topc(X21,X19))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20)))|(v2_struct_0(X19)|(~v2_pre_topc(X19)|~l1_pre_topc(X19)))))&(~r2_funct_2(u1_struct_0(X21),u1_struct_0(X20),k3_tmap_1(X19,X20,k1_tsep_1(X19,X21,X22),X21,X25),X23)|~r2_funct_2(u1_struct_0(X22),u1_struct_0(X20),k3_tmap_1(X19,X20,k1_tsep_1(X19,X21,X22),X22,X25),X24)|X25=k10_tmap_1(X19,X20,X21,X22,X23,X24)|(~v1_funct_1(X25)|(~v1_funct_2(X25,u1_struct_0(k1_tsep_1(X19,X21,X22)),u1_struct_0(X20))|~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X19,X21,X22)),u1_struct_0(X20))))))|~r1_tsep_1(X21,X22)|(~v1_funct_1(X24)|(~v1_funct_2(X24,u1_struct_0(X22),u1_struct_0(X20))|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(X20))))))|(~v1_funct_1(X23)|(~v1_funct_2(X23,u1_struct_0(X21),u1_struct_0(X20))|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X21),u1_struct_0(X20))))))|(v2_struct_0(X22)|~m1_pre_topc(X22,X19))|(v2_struct_0(X21)|~m1_pre_topc(X21,X19))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20)))|(v2_struct_0(X19)|(~v2_pre_topc(X19)|~l1_pre_topc(X19)))))&(((r2_funct_2(u1_struct_0(X21),u1_struct_0(X20),k3_tmap_1(X19,X20,k1_tsep_1(X19,X21,X22),X21,X25),X23)|X25!=k10_tmap_1(X19,X20,X21,X22,X23,X24)|(~v1_funct_1(X25)|(~v1_funct_2(X25,u1_struct_0(k1_tsep_1(X19,X21,X22)),u1_struct_0(X20))|~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X19,X21,X22)),u1_struct_0(X20))))))|~r2_funct_2(u1_struct_0(k2_tsep_1(X19,X21,X22)),u1_struct_0(X20),k3_tmap_1(X19,X20,X21,k2_tsep_1(X19,X21,X22),X23),k3_tmap_1(X19,X20,X22,k2_tsep_1(X19,X21,X22),X24))|(~v1_funct_1(X24)|(~v1_funct_2(X24,u1_struct_0(X22),u1_struct_0(X20))|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(X20))))))|(~v1_funct_1(X23)|(~v1_funct_2(X23,u1_struct_0(X21),u1_struct_0(X20))|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X21),u1_struct_0(X20))))))|(v2_struct_0(X22)|~m1_pre_topc(X22,X19))|(v2_struct_0(X21)|~m1_pre_topc(X21,X19))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20)))|(v2_struct_0(X19)|(~v2_pre_topc(X19)|~l1_pre_topc(X19))))&(r2_funct_2(u1_struct_0(X22),u1_struct_0(X20),k3_tmap_1(X19,X20,k1_tsep_1(X19,X21,X22),X22,X25),X24)|X25!=k10_tmap_1(X19,X20,X21,X22,X23,X24)|(~v1_funct_1(X25)|(~v1_funct_2(X25,u1_struct_0(k1_tsep_1(X19,X21,X22)),u1_struct_0(X20))|~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X19,X21,X22)),u1_struct_0(X20))))))|~r2_funct_2(u1_struct_0(k2_tsep_1(X19,X21,X22)),u1_struct_0(X20),k3_tmap_1(X19,X20,X21,k2_tsep_1(X19,X21,X22),X23),k3_tmap_1(X19,X20,X22,k2_tsep_1(X19,X21,X22),X24))|(~v1_funct_1(X24)|(~v1_funct_2(X24,u1_struct_0(X22),u1_struct_0(X20))|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(X20))))))|(~v1_funct_1(X23)|(~v1_funct_2(X23,u1_struct_0(X21),u1_struct_0(X20))|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X21),u1_struct_0(X20))))))|(v2_struct_0(X22)|~m1_pre_topc(X22,X19))|(v2_struct_0(X21)|~m1_pre_topc(X21,X19))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20)))|(v2_struct_0(X19)|(~v2_pre_topc(X19)|~l1_pre_topc(X19)))))&(~r2_funct_2(u1_struct_0(X21),u1_struct_0(X20),k3_tmap_1(X19,X20,k1_tsep_1(X19,X21,X22),X21,X25),X23)|~r2_funct_2(u1_struct_0(X22),u1_struct_0(X20),k3_tmap_1(X19,X20,k1_tsep_1(X19,X21,X22),X22,X25),X24)|X25=k10_tmap_1(X19,X20,X21,X22,X23,X24)|(~v1_funct_1(X25)|(~v1_funct_2(X25,u1_struct_0(k1_tsep_1(X19,X21,X22)),u1_struct_0(X20))|~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X19,X21,X22)),u1_struct_0(X20))))))|~r2_funct_2(u1_struct_0(k2_tsep_1(X19,X21,X22)),u1_struct_0(X20),k3_tmap_1(X19,X20,X21,k2_tsep_1(X19,X21,X22),X23),k3_tmap_1(X19,X20,X22,k2_tsep_1(X19,X21,X22),X24))|(~v1_funct_1(X24)|(~v1_funct_2(X24,u1_struct_0(X22),u1_struct_0(X20))|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(X20))))))|(~v1_funct_1(X23)|(~v1_funct_2(X23,u1_struct_0(X21),u1_struct_0(X20))|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X21),u1_struct_0(X20))))))|(v2_struct_0(X22)|~m1_pre_topc(X22,X19))|(v2_struct_0(X21)|~m1_pre_topc(X21,X19))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20)))|(v2_struct_0(X19)|(~v2_pre_topc(X19)|~l1_pre_topc(X19)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_tmap_1])])])])])).
fof(c_0_9, plain, ![X27, X28, X29, X30, X31, X32]:((v1_funct_1(k10_tmap_1(X27,X28,X29,X30,X31,X32))|(v2_struct_0(X27)|(~v2_pre_topc(X27)|~l1_pre_topc(X27))|(v2_struct_0(X28)|(~v2_pre_topc(X28)|~l1_pre_topc(X28))|(v2_struct_0(X29)|~m1_pre_topc(X29,X27)|(v2_struct_0(X30)|~m1_pre_topc(X30,X27)|(~v1_funct_1(X31)|(~v1_funct_2(X31,u1_struct_0(X29),u1_struct_0(X28))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X29),u1_struct_0(X28)))))|(~v1_funct_1(X32)|(~v1_funct_2(X32,u1_struct_0(X30),u1_struct_0(X28))|~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X30),u1_struct_0(X28))))))))))))&((v1_funct_2(k10_tmap_1(X27,X28,X29,X30,X31,X32),u1_struct_0(k1_tsep_1(X27,X29,X30)),u1_struct_0(X28))|(v2_struct_0(X27)|(~v2_pre_topc(X27)|~l1_pre_topc(X27))|(v2_struct_0(X28)|(~v2_pre_topc(X28)|~l1_pre_topc(X28))|(v2_struct_0(X29)|~m1_pre_topc(X29,X27)|(v2_struct_0(X30)|~m1_pre_topc(X30,X27)|(~v1_funct_1(X31)|(~v1_funct_2(X31,u1_struct_0(X29),u1_struct_0(X28))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X29),u1_struct_0(X28)))))|(~v1_funct_1(X32)|(~v1_funct_2(X32,u1_struct_0(X30),u1_struct_0(X28))|~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X30),u1_struct_0(X28))))))))))))&(m1_subset_1(k10_tmap_1(X27,X28,X29,X30,X31,X32),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X27,X29,X30)),u1_struct_0(X28))))|(v2_struct_0(X27)|(~v2_pre_topc(X27)|~l1_pre_topc(X27))|(v2_struct_0(X28)|(~v2_pre_topc(X28)|~l1_pre_topc(X28))|(v2_struct_0(X29)|~m1_pre_topc(X29,X27)|(v2_struct_0(X30)|~m1_pre_topc(X30,X27)|(~v1_funct_1(X31)|(~v1_funct_2(X31,u1_struct_0(X29),u1_struct_0(X28))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X29),u1_struct_0(X28)))))|(~v1_funct_1(X32)|(~v1_funct_2(X32,u1_struct_0(X30),u1_struct_0(X28))|~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X30),u1_struct_0(X28)))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k10_tmap_1])])])])).
fof(c_0_10, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(~(r1_tsep_1(X3,X4))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))=>((r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6))&r4_tsep_1(X1,X3,X4))=>(v1_funct_1(k10_tmap_1(X1,X2,X3,X4,X5,X6))&(v1_funct_2(k10_tmap_1(X1,X2,X3,X4,X5,X6),u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&(v5_pre_topc(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_tsep_1(X1,X3,X4),X2)&m1_subset_1(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2)))))))))))))))), inference(assume_negation,[status(cth)],[t144_tmap_1])).
cnf(c_0_11, plain, (r2_funct_2(u1_struct_0(X1),u1_struct_0(X2),k3_tmap_1(X3,X2,k1_tsep_1(X3,X4,X1),X1,X5),X6)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X3)|X5!=k10_tmap_1(X3,X2,X4,X1,X7,X6)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(k1_tsep_1(X3,X4,X1)),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X3,X4,X1)),u1_struct_0(X2))))|~r2_funct_2(u1_struct_0(k2_tsep_1(X3,X4,X1)),u1_struct_0(X2),k3_tmap_1(X3,X2,X4,k2_tsep_1(X3,X4,X1),X7),k3_tmap_1(X3,X2,X1,k2_tsep_1(X3,X4,X1),X6))|~v1_funct_1(X6)|~v1_funct_2(X6,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_1(X7)|~v1_funct_2(X7,u1_struct_0(X4),u1_struct_0(X2))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~m1_pre_topc(X1,X3)|~m1_pre_topc(X4,X3)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (v1_funct_1(k10_tmap_1(X1,X2,X3,X4,X5,X6))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_1(X6)|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (v1_funct_2(k10_tmap_1(X1,X2,X3,X4,X5,X6),u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_1(X6)|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (m1_subset_1(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_1(X6)|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_15, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk1_0))&((~v2_struct_0(esk4_0)&m1_pre_topc(esk4_0,esk1_0))&(~r1_tsep_1(esk3_0,esk4_0)&((v1_funct_1(esk5_0)&(v1_funct_2(esk5_0,u1_struct_0(esk3_0),u1_struct_0(esk2_0))&(v5_pre_topc(esk5_0,esk3_0,esk2_0)&m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0)))))))&((v1_funct_1(esk6_0)&(v1_funct_2(esk6_0,u1_struct_0(esk4_0),u1_struct_0(esk2_0))&(v5_pre_topc(esk6_0,esk4_0,esk2_0)&m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0)))))))&((r2_funct_2(u1_struct_0(k2_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,esk3_0,k2_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0),k3_tmap_1(esk1_0,esk2_0,esk4_0,k2_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0))&r4_tsep_1(esk1_0,esk3_0,esk4_0))&(~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|(~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))|(~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),esk2_0)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))))))))))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_10])])])])).
fof(c_0_16, plain, ![X42, X43, X44, X45]:((~r2_funct_2(X42,X43,X44,X45)|X44=X45|(~v1_funct_1(X44)|(~v1_funct_2(X44,X42,X43)|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(X42,X43))))|(~v1_funct_1(X45)|(~v1_funct_2(X45,X42,X43)|~m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(X42,X43)))))))&(X44!=X45|r2_funct_2(X42,X43,X44,X45)|(~v1_funct_1(X44)|(~v1_funct_2(X44,X42,X43)|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(X42,X43))))|(~v1_funct_1(X45)|(~v1_funct_2(X45,X42,X43)|~m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(X42,X43)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r2_funct_2])])])).
cnf(c_0_17, plain, (r2_funct_2(u1_struct_0(X1),u1_struct_0(X2),k3_tmap_1(X3,X2,k1_tsep_1(X3,X4,X1),X1,k10_tmap_1(X3,X2,X4,X1,X5,X6)),X6)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k2_tsep_1(X3,X4,X1)),u1_struct_0(X2),k3_tmap_1(X3,X2,X4,k2_tsep_1(X3,X4,X1),X5),k3_tmap_1(X3,X2,X1,k2_tsep_1(X3,X4,X1),X6))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_2(X6,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X5)|~v1_funct_1(X6)|~m1_pre_topc(X4,X3)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_11]), c_0_12]), c_0_13]), c_0_14])).
cnf(c_0_18, negated_conjecture, (r2_funct_2(u1_struct_0(k2_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,esk3_0,k2_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0),k3_tmap_1(esk1_0,esk2_0,esk4_0,k2_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, negated_conjecture, (m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, negated_conjecture, (v1_funct_2(esk5_0,u1_struct_0(esk3_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, negated_conjecture, (v1_funct_2(esk6_0,u1_struct_0(esk4_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, negated_conjecture, (v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, negated_conjecture, (v1_funct_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_26, negated_conjecture, (m1_pre_topc(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_27, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_28, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_29, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_30, negated_conjecture, (v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_31, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_32, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_33, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_34, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_35, plain, (r2_funct_2(u1_struct_0(X1),u1_struct_0(X2),k3_tmap_1(X3,X2,k1_tsep_1(X3,X1,X4),X1,X5),X6)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|X5!=k10_tmap_1(X3,X2,X1,X4,X6,X7)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(k1_tsep_1(X3,X1,X4)),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X3,X1,X4)),u1_struct_0(X2))))|~r2_funct_2(u1_struct_0(k2_tsep_1(X3,X1,X4)),u1_struct_0(X2),k3_tmap_1(X3,X2,X1,k2_tsep_1(X3,X1,X4),X6),k3_tmap_1(X3,X2,X4,k2_tsep_1(X3,X1,X4),X7))|~v1_funct_1(X7)|~v1_funct_2(X7,u1_struct_0(X4),u1_struct_0(X2))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_1(X6)|~v1_funct_2(X6,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_pre_topc(X4,X3)|~m1_pre_topc(X1,X3)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_36, plain, (X3=X4|~r2_funct_2(X1,X2,X3,X4)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X4)|~v1_funct_2(X4,X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_37, negated_conjecture, (r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk6_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30])]), c_0_31]), c_0_32]), c_0_33]), c_0_34])).
fof(c_0_38, plain, ![X14, X15, X16, X17, X18]:((v1_funct_1(k3_tmap_1(X14,X15,X16,X17,X18))|(v2_struct_0(X14)|(~v2_pre_topc(X14)|~l1_pre_topc(X14))|(v2_struct_0(X15)|(~v2_pre_topc(X15)|~l1_pre_topc(X15))|(~m1_pre_topc(X16,X14)|(~m1_pre_topc(X17,X14)|(~v1_funct_1(X18)|(~v1_funct_2(X18,u1_struct_0(X16),u1_struct_0(X15))|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X16),u1_struct_0(X15)))))))))))&((v1_funct_2(k3_tmap_1(X14,X15,X16,X17,X18),u1_struct_0(X17),u1_struct_0(X15))|(v2_struct_0(X14)|(~v2_pre_topc(X14)|~l1_pre_topc(X14))|(v2_struct_0(X15)|(~v2_pre_topc(X15)|~l1_pre_topc(X15))|(~m1_pre_topc(X16,X14)|(~m1_pre_topc(X17,X14)|(~v1_funct_1(X18)|(~v1_funct_2(X18,u1_struct_0(X16),u1_struct_0(X15))|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X16),u1_struct_0(X15)))))))))))&(m1_subset_1(k3_tmap_1(X14,X15,X16,X17,X18),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X17),u1_struct_0(X15))))|(v2_struct_0(X14)|(~v2_pre_topc(X14)|~l1_pre_topc(X14))|(v2_struct_0(X15)|(~v2_pre_topc(X15)|~l1_pre_topc(X15))|(~m1_pre_topc(X16,X14)|(~m1_pre_topc(X17,X14)|(~v1_funct_1(X18)|(~v1_funct_2(X18,u1_struct_0(X16),u1_struct_0(X15))|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X16),u1_struct_0(X15))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_tmap_1])])])])).
cnf(c_0_39, plain, (r2_funct_2(u1_struct_0(X1),u1_struct_0(X2),k3_tmap_1(X3,X2,k1_tsep_1(X3,X1,X4),X1,k10_tmap_1(X3,X2,X1,X4,X5,X6)),X5)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k2_tsep_1(X3,X1,X4)),u1_struct_0(X2),k3_tmap_1(X3,X2,X1,k2_tsep_1(X3,X1,X4),X5),k3_tmap_1(X3,X2,X4,k2_tsep_1(X3,X1,X4),X6))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_pre_topc(X4,X3)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_35]), c_0_12]), c_0_13]), c_0_14])).
cnf(c_0_40, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk6_0|~m1_subset_1(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~v1_funct_2(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_20]), c_0_22]), c_0_24])])).
cnf(c_0_41, plain, (m1_subset_1(k3_tmap_1(X1,X2,X3,X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_42, negated_conjecture, (r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk5_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_18]), c_0_20]), c_0_19]), c_0_22]), c_0_21]), c_0_24]), c_0_23]), c_0_26]), c_0_25]), c_0_27]), c_0_28]), c_0_29]), c_0_30])]), c_0_34]), c_0_32]), c_0_33]), c_0_31])).
cnf(c_0_43, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk6_0|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))))|~v1_funct_2(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))|~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_26]), c_0_28]), c_0_27]), c_0_30]), c_0_29])]), c_0_32]), c_0_33])).
cnf(c_0_44, plain, (v1_funct_2(k3_tmap_1(X1,X2,X3,X4,X5),u1_struct_0(X4),u1_struct_0(X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_45, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk5_0|~m1_subset_1(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))|~v1_funct_2(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_42]), c_0_19]), c_0_21]), c_0_23])])).
fof(c_0_46, plain, ![X5, X4, X3, X2, X1]:(epred1_5(X1,X2,X3,X4,X5)<=>((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&(v5_pre_topc(X5,k1_tsep_1(X1,X3,X4),X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2)))))))<=>((v1_funct_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5))&(v1_funct_2(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),X3,X2)&m1_subset_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))&(v1_funct_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5))&(v1_funct_2(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5),u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5),X4,X2)&m1_subset_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))))), introduced(definition)).
cnf(c_0_47, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk6_0|~v1_funct_2(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))|~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_14]), c_0_20]), c_0_19]), c_0_22]), c_0_21]), c_0_24]), c_0_23]), c_0_26]), c_0_25]), c_0_28]), c_0_27]), c_0_30]), c_0_29])]), c_0_31]), c_0_34]), c_0_32]), c_0_33])).
cnf(c_0_48, plain, (v1_funct_2(k3_tmap_1(X1,X2,k1_tsep_1(X3,X4,X5),X6,k10_tmap_1(X3,X2,X4,X5,X7,X8)),u1_struct_0(X6),u1_struct_0(X2))|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X2))))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X8,u1_struct_0(X5),u1_struct_0(X2))|~v1_funct_2(X7,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X8)|~v1_funct_1(X7)|~m1_pre_topc(k1_tsep_1(X3,X4,X5),X1)|~m1_pre_topc(X6,X1)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~l1_pre_topc(X3)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_14]), c_0_12]), c_0_13])).
cnf(c_0_49, plain, (v1_funct_1(k3_tmap_1(X1,X2,X3,X4,X5))|v2_struct_0(X1)|v2_struct_0(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk5_0|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))))|~v1_funct_2(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))|~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_41]), c_0_25]), c_0_28]), c_0_27]), c_0_30]), c_0_29])]), c_0_32]), c_0_33])).
fof(c_0_51, plain, ![X5, X4, X3, X2, X1]:(epred1_5(X1,X2,X3,X4,X5)=>((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&(v5_pre_topc(X5,k1_tsep_1(X1,X3,X4),X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2)))))))<=>((v1_funct_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5))&(v1_funct_2(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),X3,X2)&m1_subset_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))&(v1_funct_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5))&(v1_funct_2(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5),u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5),X4,X2)&m1_subset_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))))), inference(split_equiv,[status(thm)],[c_0_46])).
cnf(c_0_52, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk6_0|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))|~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_20]), c_0_19]), c_0_22]), c_0_21]), c_0_24]), c_0_23]), c_0_26]), c_0_25]), c_0_28]), c_0_27]), c_0_30]), c_0_29])]), c_0_33]), c_0_34]), c_0_31]), c_0_32])).
cnf(c_0_53, plain, (v1_funct_1(k3_tmap_1(X1,X2,k1_tsep_1(X3,X4,X5),X6,k10_tmap_1(X3,X2,X4,X5,X7,X8)))|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X2))))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X8,u1_struct_0(X5),u1_struct_0(X2))|~v1_funct_2(X7,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X8)|~v1_funct_1(X7)|~m1_pre_topc(k1_tsep_1(X3,X4,X5),X1)|~m1_pre_topc(X6,X1)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~l1_pre_topc(X3)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_14]), c_0_12]), c_0_13])).
cnf(c_0_54, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk5_0|~v1_funct_2(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))|~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_14]), c_0_20]), c_0_19]), c_0_22]), c_0_21]), c_0_24]), c_0_23]), c_0_26]), c_0_25]), c_0_28]), c_0_27]), c_0_30]), c_0_29])]), c_0_31]), c_0_34]), c_0_32]), c_0_33])).
fof(c_0_55, plain, ![X46, X47, X48, X49, X50]:((((v1_funct_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46))|(~v1_funct_1(X46)|(~v1_funct_2(X46,u1_struct_0(k1_tsep_1(X50,X48,X47)),u1_struct_0(X49))|(~v5_pre_topc(X46,k1_tsep_1(X50,X48,X47),X49)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X50,X48,X47)),u1_struct_0(X49)))))))|~epred1_5(X50,X49,X48,X47,X46))&((v1_funct_2(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46),u1_struct_0(X48),u1_struct_0(X49))|(~v1_funct_1(X46)|(~v1_funct_2(X46,u1_struct_0(k1_tsep_1(X50,X48,X47)),u1_struct_0(X49))|(~v5_pre_topc(X46,k1_tsep_1(X50,X48,X47),X49)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X50,X48,X47)),u1_struct_0(X49)))))))|~epred1_5(X50,X49,X48,X47,X46))&((v5_pre_topc(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46),X48,X49)|(~v1_funct_1(X46)|(~v1_funct_2(X46,u1_struct_0(k1_tsep_1(X50,X48,X47)),u1_struct_0(X49))|(~v5_pre_topc(X46,k1_tsep_1(X50,X48,X47),X49)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X50,X48,X47)),u1_struct_0(X49)))))))|~epred1_5(X50,X49,X48,X47,X46))&(m1_subset_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X48),u1_struct_0(X49))))|(~v1_funct_1(X46)|(~v1_funct_2(X46,u1_struct_0(k1_tsep_1(X50,X48,X47)),u1_struct_0(X49))|(~v5_pre_topc(X46,k1_tsep_1(X50,X48,X47),X49)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X50,X48,X47)),u1_struct_0(X49)))))))|~epred1_5(X50,X49,X48,X47,X46)))))&((v1_funct_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46))|(~v1_funct_1(X46)|(~v1_funct_2(X46,u1_struct_0(k1_tsep_1(X50,X48,X47)),u1_struct_0(X49))|(~v5_pre_topc(X46,k1_tsep_1(X50,X48,X47),X49)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X50,X48,X47)),u1_struct_0(X49)))))))|~epred1_5(X50,X49,X48,X47,X46))&((v1_funct_2(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46),u1_struct_0(X47),u1_struct_0(X49))|(~v1_funct_1(X46)|(~v1_funct_2(X46,u1_struct_0(k1_tsep_1(X50,X48,X47)),u1_struct_0(X49))|(~v5_pre_topc(X46,k1_tsep_1(X50,X48,X47),X49)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X50,X48,X47)),u1_struct_0(X49)))))))|~epred1_5(X50,X49,X48,X47,X46))&((v5_pre_topc(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46),X47,X49)|(~v1_funct_1(X46)|(~v1_funct_2(X46,u1_struct_0(k1_tsep_1(X50,X48,X47)),u1_struct_0(X49))|(~v5_pre_topc(X46,k1_tsep_1(X50,X48,X47),X49)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X50,X48,X47)),u1_struct_0(X49)))))))|~epred1_5(X50,X49,X48,X47,X46))&(m1_subset_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X47),u1_struct_0(X49))))|(~v1_funct_1(X46)|(~v1_funct_2(X46,u1_struct_0(k1_tsep_1(X50,X48,X47)),u1_struct_0(X49))|(~v5_pre_topc(X46,k1_tsep_1(X50,X48,X47),X49)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X50,X48,X47)),u1_struct_0(X49)))))))|~epred1_5(X50,X49,X48,X47,X46))))))&((v1_funct_1(X46)|(~v1_funct_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46))|(~v1_funct_2(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46),u1_struct_0(X48),u1_struct_0(X49))|(~v5_pre_topc(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46),X48,X49)|~m1_subset_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X48),u1_struct_0(X49))))))|(~v1_funct_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46))|(~v1_funct_2(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46),u1_struct_0(X47),u1_struct_0(X49))|(~v5_pre_topc(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46),X47,X49)|~m1_subset_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X47),u1_struct_0(X49))))))))|~epred1_5(X50,X49,X48,X47,X46))&((v1_funct_2(X46,u1_struct_0(k1_tsep_1(X50,X48,X47)),u1_struct_0(X49))|(~v1_funct_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46))|(~v1_funct_2(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46),u1_struct_0(X48),u1_struct_0(X49))|(~v5_pre_topc(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46),X48,X49)|~m1_subset_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X48),u1_struct_0(X49))))))|(~v1_funct_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46))|(~v1_funct_2(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46),u1_struct_0(X47),u1_struct_0(X49))|(~v5_pre_topc(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46),X47,X49)|~m1_subset_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X47),u1_struct_0(X49))))))))|~epred1_5(X50,X49,X48,X47,X46))&((v5_pre_topc(X46,k1_tsep_1(X50,X48,X47),X49)|(~v1_funct_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46))|(~v1_funct_2(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46),u1_struct_0(X48),u1_struct_0(X49))|(~v5_pre_topc(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46),X48,X49)|~m1_subset_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X48),u1_struct_0(X49))))))|(~v1_funct_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46))|(~v1_funct_2(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46),u1_struct_0(X47),u1_struct_0(X49))|(~v5_pre_topc(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46),X47,X49)|~m1_subset_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X47),u1_struct_0(X49))))))))|~epred1_5(X50,X49,X48,X47,X46))&(m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X50,X48,X47)),u1_struct_0(X49))))|(~v1_funct_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46))|(~v1_funct_2(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46),u1_struct_0(X48),u1_struct_0(X49))|(~v5_pre_topc(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46),X48,X49)|~m1_subset_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X48,X46),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X48),u1_struct_0(X49))))))|(~v1_funct_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46))|(~v1_funct_2(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46),u1_struct_0(X47),u1_struct_0(X49))|(~v5_pre_topc(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46),X47,X49)|~m1_subset_1(k3_tmap_1(X50,X49,k1_tsep_1(X50,X48,X47),X47,X46),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X47),u1_struct_0(X49))))))))|~epred1_5(X50,X49,X48,X47,X46)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_51])])])).
cnf(c_0_56, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk6_0|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_20]), c_0_19]), c_0_22]), c_0_21]), c_0_24]), c_0_23]), c_0_26]), c_0_25]), c_0_28]), c_0_27]), c_0_30]), c_0_29])]), c_0_33]), c_0_34]), c_0_31]), c_0_32])).
cnf(c_0_57, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk5_0|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))|~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_48]), c_0_20]), c_0_19]), c_0_22]), c_0_21]), c_0_24]), c_0_23]), c_0_25]), c_0_26]), c_0_28]), c_0_27]), c_0_30]), c_0_29])]), c_0_33]), c_0_34]), c_0_31]), c_0_32])).
fof(c_0_58, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(r4_tsep_1(X1,X3,X4)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))))=>epred1_5(X1,X2,X3,X4,X5))))))), inference(apply_def,[status(thm)],[t126_tmap_1, c_0_46])).
cnf(c_0_59, plain, (v5_pre_topc(X1,k1_tsep_1(X2,X3,X4),X5)|~v1_funct_1(k3_tmap_1(X2,X5,k1_tsep_1(X2,X3,X4),X3,X1))|~v1_funct_2(k3_tmap_1(X2,X5,k1_tsep_1(X2,X3,X4),X3,X1),u1_struct_0(X3),u1_struct_0(X5))|~v5_pre_topc(k3_tmap_1(X2,X5,k1_tsep_1(X2,X3,X4),X3,X1),X3,X5)|~m1_subset_1(k3_tmap_1(X2,X5,k1_tsep_1(X2,X3,X4),X3,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X5))))|~v1_funct_1(k3_tmap_1(X2,X5,k1_tsep_1(X2,X3,X4),X4,X1))|~v1_funct_2(k3_tmap_1(X2,X5,k1_tsep_1(X2,X3,X4),X4,X1),u1_struct_0(X4),u1_struct_0(X5))|~v5_pre_topc(k3_tmap_1(X2,X5,k1_tsep_1(X2,X3,X4),X4,X1),X4,X5)|~m1_subset_1(k3_tmap_1(X2,X5,k1_tsep_1(X2,X3,X4),X4,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X5))))|~epred1_5(X2,X5,X3,X4,X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_60, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk6_0|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_13]), c_0_20]), c_0_19]), c_0_22]), c_0_21]), c_0_24]), c_0_23]), c_0_26]), c_0_25]), c_0_28]), c_0_27]), c_0_30]), c_0_29])]), c_0_31]), c_0_34]), c_0_32]), c_0_33])).
cnf(c_0_61, negated_conjecture, (v5_pre_topc(esk6_0,esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_62, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk5_0|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_53]), c_0_20]), c_0_19]), c_0_22]), c_0_21]), c_0_24]), c_0_23]), c_0_25]), c_0_26]), c_0_28]), c_0_27]), c_0_30]), c_0_29])]), c_0_33]), c_0_34]), c_0_31]), c_0_32])).
fof(c_0_63, plain, ![X33, X34, X35, X36, X37]:(v2_struct_0(X33)|(~v2_pre_topc(X33)|~l1_pre_topc(X33))|(v2_struct_0(X34)|(~v2_pre_topc(X34)|~l1_pre_topc(X34))|(v2_struct_0(X35)|~m1_pre_topc(X35,X33)|(v2_struct_0(X36)|~m1_pre_topc(X36,X33)|(~r4_tsep_1(X33,X35,X36)|(~v1_funct_1(X37)|(~v1_funct_2(X37,u1_struct_0(k1_tsep_1(X33,X35,X36)),u1_struct_0(X34))|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X33,X35,X36)),u1_struct_0(X34)))))|epred1_5(X33,X34,X35,X36,X37))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_58])])])])).
cnf(c_0_64, plain, (v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),esk2_0)|~epred1_5(esk1_0,esk2_0,esk3_0,esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_subset_1(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))|~v5_pre_topc(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk3_0,esk2_0)|~v1_funct_2(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_20]), c_0_61]), c_0_22]), c_0_24])])).
cnf(c_0_65, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk5_0|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_13]), c_0_20]), c_0_19]), c_0_22]), c_0_21]), c_0_24]), c_0_23]), c_0_26]), c_0_25]), c_0_28]), c_0_27]), c_0_30]), c_0_29])]), c_0_31]), c_0_34]), c_0_32]), c_0_33])).
cnf(c_0_66, negated_conjecture, (v5_pre_topc(esk5_0,esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_67, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|epred1_5(X1,X2,X3,X4,X5)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~r4_tsep_1(X1,X3,X4)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_68, negated_conjecture, (~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),esk2_0)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_69, negated_conjecture, (v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),esk2_0)|~epred1_5(esk1_0,esk2_0,esk3_0,esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_19]), c_0_66]), c_0_21]), c_0_23])])).
cnf(c_0_70, plain, (epred1_5(X1,X2,X3,X4,k10_tmap_1(X1,X2,X3,X4,X5,X6))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~r4_tsep_1(X1,X3,X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_14]), c_0_12]), c_0_13])).
cnf(c_0_71, negated_conjecture, (r4_tsep_1(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_72, negated_conjecture, (~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),esk2_0)|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_14]), c_0_20]), c_0_19]), c_0_22]), c_0_21]), c_0_24]), c_0_23]), c_0_26]), c_0_25]), c_0_28]), c_0_27]), c_0_30]), c_0_29])]), c_0_31]), c_0_34]), c_0_32]), c_0_33])).
cnf(c_0_73, negated_conjecture, (v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),esk2_0)|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_71]), c_0_20]), c_0_19]), c_0_22]), c_0_21]), c_0_24]), c_0_23]), c_0_26]), c_0_25]), c_0_28]), c_0_27]), c_0_30]), c_0_29])]), c_0_33]), c_0_32]), c_0_34]), c_0_31])).
cnf(c_0_74, negated_conjecture, (~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk1_0)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
fof(c_0_75, plain, ![X26]:(~l1_pre_topc(X26)|m1_pre_topc(X26,X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tsep_1])])).
cnf(c_0_76, negated_conjecture, (~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_13]), c_0_20]), c_0_19]), c_0_22]), c_0_21]), c_0_24]), c_0_23]), c_0_26]), c_0_25]), c_0_28]), c_0_27]), c_0_30]), c_0_29])]), c_0_31]), c_0_34]), c_0_32]), c_0_33])).
fof(c_0_77, plain, ![X38, X39, X40, X41]:(v2_struct_0(X38)|(~v2_pre_topc(X38)|~l1_pre_topc(X38))|(v2_struct_0(X39)|~m1_pre_topc(X39,X38)|(v2_struct_0(X40)|~m1_pre_topc(X40,X38)|(v2_struct_0(X41)|~m1_pre_topc(X41,X38)|(~m1_pre_topc(X39,X40)|~m1_pre_topc(X41,X40)|m1_pre_topc(k1_tsep_1(X38,X39,X41),X40)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t29_tmap_1])])])])).
cnf(c_0_78, plain, (m1_pre_topc(X1,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_79, negated_conjecture, (~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_12]), c_0_20]), c_0_19]), c_0_22]), c_0_21]), c_0_24]), c_0_23]), c_0_26]), c_0_25]), c_0_28]), c_0_27]), c_0_30]), c_0_29])]), c_0_31]), c_0_34]), c_0_32]), c_0_33])).
cnf(c_0_80, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|m1_pre_topc(k1_tsep_1(X1,X2,X4),X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_81, negated_conjecture, (m1_pre_topc(esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_78, c_0_27])).
cnf(c_0_82, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_26]), c_0_81]), c_0_25]), c_0_27]), c_0_29])]), c_0_31]), c_0_33]), c_0_34]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 83
# Proof object clause steps            : 62
# Proof object formula steps           : 21
# Proof object conjectures             : 46
# Proof object clause conjectures      : 43
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 34
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 26
# Proof object simplifying inferences  : 296
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 51
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 51
# Processed clauses                    : 196
# ...of these trivial                  : 0
# ...subsumed                          : 14
# ...remaining for further processing  : 182
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 29
# Backward-rewritten                   : 0
# Generated clauses                    : 158
# ...of the previous two non-trivial   : 133
# Contextual simplify-reflections      : 55
# Paramodulations                      : 153
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 97
#    Positive orientable unit clauses  : 20
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 71
# Current number of unprocessed clauses: 39
# ...number of literals in the above   : 676
# Current number of archived formulas  : 0
# Current number of archived clauses   : 80
# Clause-clause subsumption calls (NU) : 10936
# Rec. Clause-clause subsumption calls : 595
# Non-unit clause-clause subsumptions  : 94
# Unit Clause-clause subsumption calls : 10
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 24922

# -------------------------------------------------
# User time                : 0.031 s
# System time              : 0.000 s
# Total time               : 0.031 s
# Maximum resident set size: 28816 pages
