# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/1c6a18bac55a46388e84077b9ba6e4c8.E_input', redefinition_k6_domain_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/1c6a18bac55a46388e84077b9ba6e4c8.E_input', t69_enumset1)).
fof(t38_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))&r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))))), file('data/MPTP/proofs/1c6a18bac55a46388e84077b9ba6e4c8.E_input', t38_yellow_0)).
fof(t7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2)))))))), file('data/MPTP/proofs/1c6a18bac55a46388e84077b9ba6e4c8.E_input', t7_yellow_0)).
fof(t31_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3))=>(r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2)))))&((r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2))))=>(X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3)))))), file('data/MPTP/proofs/1c6a18bac55a46388e84077b9ba6e4c8.E_input', t31_yellow_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/1c6a18bac55a46388e84077b9ba6e4c8.E_input', fc2_struct_0)).
fof(t30_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4)))))&((r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))=>(X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3)))))), file('data/MPTP/proofs/1c6a18bac55a46388e84077b9ba6e4c8.E_input', t30_yellow_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/1c6a18bac55a46388e84077b9ba6e4c8.E_input', dt_l1_orders_2)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/1c6a18bac55a46388e84077b9ba6e4c8.E_input', t24_orders_2)).
fof(c_0_9, plain, ![X17, X18]:(v1_xboole_0(X17)|~m1_subset_1(X18,X17)|k6_domain_1(X17,X18)=k1_tarski(X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_10, plain, ![X19]:k2_tarski(X19,X19)=k1_tarski(X19), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_11, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))&r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2)))))), inference(assume_negation,[status(cth)],[t38_yellow_0])).
fof(c_0_12, plain, ![X8, X9, X10]:((~r1_lattice3(X8,k1_tarski(X10),X9)|r1_orders_2(X8,X9,X10)|~m1_subset_1(X10,u1_struct_0(X8))|~m1_subset_1(X9,u1_struct_0(X8))|~l1_orders_2(X8))&((~r1_orders_2(X8,X9,X10)|r1_lattice3(X8,k1_tarski(X10),X9)|~m1_subset_1(X10,u1_struct_0(X8))|~m1_subset_1(X9,u1_struct_0(X8))|~l1_orders_2(X8))&((~r2_lattice3(X8,k1_tarski(X10),X9)|r1_orders_2(X8,X10,X9)|~m1_subset_1(X10,u1_struct_0(X8))|~m1_subset_1(X9,u1_struct_0(X8))|~l1_orders_2(X8))&(~r1_orders_2(X8,X10,X9)|r2_lattice3(X8,k1_tarski(X10),X9)|~m1_subset_1(X10,u1_struct_0(X8))|~m1_subset_1(X9,u1_struct_0(X8))|~l1_orders_2(X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])).
cnf(c_0_13, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(~r1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))|~r2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])).
cnf(c_0_16, plain, (r1_orders_2(X1,X3,X2)|~r1_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_17, plain, ![X23, X24, X25, X26, X27]:(((r1_lattice3(X23,X25,X24)|(X24!=k2_yellow_0(X23,X25)|~r2_yellow_0(X23,X25))|~m1_subset_1(X24,u1_struct_0(X23))|(~v5_orders_2(X23)|~l1_orders_2(X23)))&(~m1_subset_1(X26,u1_struct_0(X23))|(~r1_lattice3(X23,X25,X26)|r1_orders_2(X23,X26,X24))|(X24!=k2_yellow_0(X23,X25)|~r2_yellow_0(X23,X25))|~m1_subset_1(X24,u1_struct_0(X23))|(~v5_orders_2(X23)|~l1_orders_2(X23))))&(((X24=k2_yellow_0(X23,X27)|(m1_subset_1(esk4_3(X23,X24,X27),u1_struct_0(X23))|~r1_lattice3(X23,X27,X24))|~m1_subset_1(X24,u1_struct_0(X23))|(~v5_orders_2(X23)|~l1_orders_2(X23)))&(r2_yellow_0(X23,X27)|(m1_subset_1(esk4_3(X23,X24,X27),u1_struct_0(X23))|~r1_lattice3(X23,X27,X24))|~m1_subset_1(X24,u1_struct_0(X23))|(~v5_orders_2(X23)|~l1_orders_2(X23))))&(((X24=k2_yellow_0(X23,X27)|(r1_lattice3(X23,X27,esk4_3(X23,X24,X27))|~r1_lattice3(X23,X27,X24))|~m1_subset_1(X24,u1_struct_0(X23))|(~v5_orders_2(X23)|~l1_orders_2(X23)))&(r2_yellow_0(X23,X27)|(r1_lattice3(X23,X27,esk4_3(X23,X24,X27))|~r1_lattice3(X23,X27,X24))|~m1_subset_1(X24,u1_struct_0(X23))|(~v5_orders_2(X23)|~l1_orders_2(X23))))&((X24=k2_yellow_0(X23,X27)|(~r1_orders_2(X23,esk4_3(X23,X24,X27),X24)|~r1_lattice3(X23,X27,X24))|~m1_subset_1(X24,u1_struct_0(X23))|(~v5_orders_2(X23)|~l1_orders_2(X23)))&(r2_yellow_0(X23,X27)|(~r1_orders_2(X23,esk4_3(X23,X24,X27),X24)|~r1_lattice3(X23,X27,X24))|~m1_subset_1(X24,u1_struct_0(X23))|(~v5_orders_2(X23)|~l1_orders_2(X23))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_yellow_0])])])])])])).
fof(c_0_18, plain, ![X7]:(v2_struct_0(X7)|~l1_struct_0(X7)|~v1_xboole_0(u1_struct_0(X7))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_19, plain, (k2_tarski(X2,X2)=k6_domain_1(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_20, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,k2_tarski(X2,X2),X3)), inference(rw,[status(thm)],[c_0_16, c_0_14])).
cnf(c_0_22, plain, (r2_yellow_0(X1,X2)|r1_lattice3(X1,X2,esk4_3(X1,X3,X2))|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (r2_yellow_0(X1,X2)|m1_subset_1(esk4_3(X1,X3,X2),u1_struct_0(X1))|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (r1_orders_2(X1,X2,X3)|~r2_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_25, plain, ![X11, X12, X13, X14, X15]:(((r2_lattice3(X11,X13,X12)|(X12!=k1_yellow_0(X11,X13)|~r1_yellow_0(X11,X13))|~m1_subset_1(X12,u1_struct_0(X11))|(~v5_orders_2(X11)|~l1_orders_2(X11)))&(~m1_subset_1(X14,u1_struct_0(X11))|(~r2_lattice3(X11,X13,X14)|r1_orders_2(X11,X12,X14))|(X12!=k1_yellow_0(X11,X13)|~r1_yellow_0(X11,X13))|~m1_subset_1(X12,u1_struct_0(X11))|(~v5_orders_2(X11)|~l1_orders_2(X11))))&(((X12=k1_yellow_0(X11,X15)|(m1_subset_1(esk3_3(X11,X12,X15),u1_struct_0(X11))|~r2_lattice3(X11,X15,X12))|~m1_subset_1(X12,u1_struct_0(X11))|(~v5_orders_2(X11)|~l1_orders_2(X11)))&(r1_yellow_0(X11,X15)|(m1_subset_1(esk3_3(X11,X12,X15),u1_struct_0(X11))|~r2_lattice3(X11,X15,X12))|~m1_subset_1(X12,u1_struct_0(X11))|(~v5_orders_2(X11)|~l1_orders_2(X11))))&(((X12=k1_yellow_0(X11,X15)|(r2_lattice3(X11,X15,esk3_3(X11,X12,X15))|~r2_lattice3(X11,X15,X12))|~m1_subset_1(X12,u1_struct_0(X11))|(~v5_orders_2(X11)|~l1_orders_2(X11)))&(r1_yellow_0(X11,X15)|(r2_lattice3(X11,X15,esk3_3(X11,X12,X15))|~r2_lattice3(X11,X15,X12))|~m1_subset_1(X12,u1_struct_0(X11))|(~v5_orders_2(X11)|~l1_orders_2(X11))))&((X12=k1_yellow_0(X11,X15)|(~r1_orders_2(X11,X12,esk3_3(X11,X12,X15))|~r2_lattice3(X11,X15,X12))|~m1_subset_1(X12,u1_struct_0(X11))|(~v5_orders_2(X11)|~l1_orders_2(X11)))&(r1_yellow_0(X11,X15)|(~r1_orders_2(X11,X12,esk3_3(X11,X12,X15))|~r2_lattice3(X11,X15,X12))|~m1_subset_1(X12,u1_struct_0(X11))|(~v5_orders_2(X11)|~l1_orders_2(X11))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_yellow_0])])])])])])).
cnf(c_0_26, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_28, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_29, plain, (r2_yellow_0(X1,X2)|~r1_orders_2(X1,esk4_3(X1,X3,X2),X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_30, plain, (r1_orders_2(X1,esk4_3(X1,X2,k2_tarski(X3,X3)),X3)|r2_yellow_0(X1,k2_tarski(X3,X3))|~r1_lattice3(X1,k2_tarski(X3,X3),X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])).
cnf(c_0_31, plain, (r1_lattice3(X1,k1_tarski(X3),X2)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_32, plain, (r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r2_lattice3(X1,k2_tarski(X2,X2),X3)), inference(rw,[status(thm)],[c_0_24, c_0_14])).
cnf(c_0_33, plain, (r1_yellow_0(X1,X2)|r2_lattice3(X1,X2,esk3_3(X1,X3,X2))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, plain, (r1_yellow_0(X1,X2)|m1_subset_1(esk3_3(X1,X3,X2),u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, negated_conjecture, (~r1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))|~r2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_36, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])).
cnf(c_0_37, plain, (r2_yellow_0(X1,k2_tarski(X2,X2))|~r1_lattice3(X1,k2_tarski(X2,X2),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_38, plain, (r1_lattice3(X1,k2_tarski(X3,X3),X2)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[c_0_31, c_0_14])).
cnf(c_0_39, plain, (r1_yellow_0(X1,X2)|~r1_orders_2(X1,X3,esk3_3(X1,X3,X2))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_40, plain, (r1_orders_2(X1,X2,esk3_3(X1,X3,k2_tarski(X2,X2)))|r1_yellow_0(X1,k2_tarski(X2,X2))|~r2_lattice3(X1,k2_tarski(X2,X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])).
cnf(c_0_41, plain, (r2_lattice3(X1,k1_tarski(X2),X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_42, negated_conjecture, (~l1_struct_0(esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~r1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_43, plain, (r2_yellow_0(X1,k2_tarski(X2,X2))|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_44, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_45, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_46, plain, (r1_yellow_0(X1,k2_tarski(X2,X2))|~r2_lattice3(X1,k2_tarski(X2,X2),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_47, plain, (r2_lattice3(X1,k2_tarski(X2,X2),X3)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[c_0_41, c_0_14])).
cnf(c_0_48, negated_conjecture, (~r1_orders_2(esk1_0,esk2_0,esk2_0)|~l1_struct_0(esk1_0)|~r1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_20]), c_0_44]), c_0_45])])).
cnf(c_0_49, plain, (r1_yellow_0(X1,k2_tarski(X2,X2))|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
fof(c_0_50, plain, ![X22]:(~l1_orders_2(X22)|l1_struct_0(X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
fof(c_0_51, plain, ![X20, X21]:(v2_struct_0(X20)|(~v3_orders_2(X20)|~l1_orders_2(X20))|(~m1_subset_1(X21,u1_struct_0(X20))|r1_orders_2(X20,X21,X21))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
cnf(c_0_52, negated_conjecture, (~r1_orders_2(esk1_0,esk2_0,esk2_0)|~l1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_20]), c_0_44]), c_0_45])])).
cnf(c_0_53, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_54, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_55, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_56, negated_conjecture, (~r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_44])])).
cnf(c_0_57, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_44])]), c_0_28])).
cnf(c_0_58, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_20])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 59
# Proof object clause steps            : 40
# Proof object formula steps           : 19
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 23
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 30
# Processed clauses                    : 88
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 88
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 94
# ...of the previous two non-trivial   : 88
# Contextual simplify-reflections      : 7
# Paramodulations                      : 90
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 53
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 47
# Current number of unprocessed clauses: 60
# ...number of literals in the above   : 673
# Current number of archived formulas  : 0
# Current number of archived clauses   : 32
# Clause-clause subsumption calls (NU) : 1302
# Rec. Clause-clause subsumption calls : 38
# Non-unit clause-clause subsumptions  : 8
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6128

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.000 s
# Total time               : 0.013 s
# Maximum resident set size: 28700 pages
