# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', d5_xboole_0)).
fof(t52_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(X1,k4_xboole_0(X2,X3))=k2_xboole_0(k4_xboole_0(X1,X2),k3_xboole_0(X1,X3)), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', t52_xboole_1)).
fof(t4_boole, axiom, ![X1]:k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0, file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', t4_boole)).
fof(d3_xboole_0, axiom, ![X1, X2, X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2)))), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', d3_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', commutativity_k3_xboole_0)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', t1_boole)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', t7_xboole_0)).
fof(t75_xboole_1, axiom, ![X1, X2]:~((~(r1_xboole_0(X1,X2))&r1_xboole_0(k3_xboole_0(X1,X2),X2))), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', t75_xboole_1)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', t83_xboole_1)).
fof(t68_zfmisc_1, conjecture, ![X1, X2]:(k4_xboole_0(k1_tarski(X1),X2)=k1_xboole_0<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', t68_zfmisc_1)).
fof(commutativity_k5_xboole_0, axiom, ![X1, X2]:k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', commutativity_k5_xboole_0)).
fof(t101_xboole_1, axiom, ![X1, X2]:k5_xboole_0(X1,X2)=k4_xboole_0(k2_xboole_0(X1,X2),k3_xboole_0(X1,X2)), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', t101_xboole_1)).
fof(t76_enumset1, axiom, ![X1]:k1_enumset1(X1,X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', t76_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', t71_enumset1)).
fof(t45_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X4)), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', t45_enumset1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', d1_tarski)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', symmetry_r1_xboole_0)).
fof(idempotence_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', idempotence_k2_xboole_0)).
fof(t50_zfmisc_1, axiom, ![X1, X2, X3]:~(k2_xboole_0(k2_tarski(X1,X2),X3)=k1_xboole_0), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', t50_zfmisc_1)).
fof(t50_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3)), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', t50_xboole_1)).
fof(t49_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),X3), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', t49_xboole_1)).
fof(t88_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>k4_xboole_0(k2_xboole_0(X1,X2),X2)=X1), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', t88_xboole_1)).
fof(t24_xboole_1, axiom, ![X1, X2, X3]:k2_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(k2_xboole_0(X1,X2),k2_xboole_0(X1,X3)), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', t24_xboole_1)).
fof(t48_zfmisc_1, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&r2_hidden(X3,X2))=>k2_xboole_0(k2_tarski(X1,X3),X2)=X2), file('data/MPTP/proofs/1c70a2aea79a4dc493a3450154646a78.E_input', t48_zfmisc_1)).
fof(c_0_24, plain, ![X23, X24, X25, X26, X27, X28, X29, X30]:((((r2_hidden(X26,X23)|~r2_hidden(X26,X25)|X25!=k4_xboole_0(X23,X24))&(~r2_hidden(X26,X24)|~r2_hidden(X26,X25)|X25!=k4_xboole_0(X23,X24)))&(~r2_hidden(X27,X23)|r2_hidden(X27,X24)|r2_hidden(X27,X25)|X25!=k4_xboole_0(X23,X24)))&((~r2_hidden(esk4_3(X28,X29,X30),X30)|(~r2_hidden(esk4_3(X28,X29,X30),X28)|r2_hidden(esk4_3(X28,X29,X30),X29))|X30=k4_xboole_0(X28,X29))&((r2_hidden(esk4_3(X28,X29,X30),X28)|r2_hidden(esk4_3(X28,X29,X30),X30)|X30=k4_xboole_0(X28,X29))&(~r2_hidden(esk4_3(X28,X29,X30),X29)|r2_hidden(esk4_3(X28,X29,X30),X30)|X30=k4_xboole_0(X28,X29))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_25, plain, ![X52, X53, X54]:k4_xboole_0(X52,k4_xboole_0(X53,X54))=k2_xboole_0(k4_xboole_0(X52,X53),k3_xboole_0(X52,X54)), inference(variable_rename,[status(thm)],[t52_xboole_1])).
fof(c_0_26, plain, ![X22]:k4_xboole_0(k1_xboole_0,X22)=k1_xboole_0, inference(variable_rename,[status(thm)],[t4_boole])).
fof(c_0_27, plain, ![X32, X33, X34, X35, X36, X37, X38, X39]:(((~r2_hidden(X35,X34)|(r2_hidden(X35,X32)|r2_hidden(X35,X33))|X34!=k2_xboole_0(X32,X33))&((~r2_hidden(X36,X32)|r2_hidden(X36,X34)|X34!=k2_xboole_0(X32,X33))&(~r2_hidden(X36,X33)|r2_hidden(X36,X34)|X34!=k2_xboole_0(X32,X33))))&(((~r2_hidden(esk5_3(X37,X38,X39),X37)|~r2_hidden(esk5_3(X37,X38,X39),X39)|X39=k2_xboole_0(X37,X38))&(~r2_hidden(esk5_3(X37,X38,X39),X38)|~r2_hidden(esk5_3(X37,X38,X39),X39)|X39=k2_xboole_0(X37,X38)))&(r2_hidden(esk5_3(X37,X38,X39),X39)|(r2_hidden(esk5_3(X37,X38,X39),X37)|r2_hidden(esk5_3(X37,X38,X39),X38))|X39=k2_xboole_0(X37,X38)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])).
cnf(c_0_28, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X3))=k2_xboole_0(k4_xboole_0(X1,X2),k3_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_31, plain, ![X60, X61]:k3_xboole_0(X60,X61)=k3_xboole_0(X61,X60), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_32, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_33, plain, ![X41]:k2_xboole_0(X41,k1_xboole_0)=X41, inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_34, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_28])).
cnf(c_0_35, plain, (k2_xboole_0(k1_xboole_0,k3_xboole_0(k1_xboole_0,X1))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_30])).
cnf(c_0_36, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_37, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_38, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_39, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_34, c_0_30])).
cnf(c_0_40, plain, (k2_xboole_0(k1_xboole_0,k3_xboole_0(X1,k1_xboole_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_41, plain, (~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
fof(c_0_42, plain, ![X47]:(X47=k1_xboole_0|r2_hidden(esk6_1(X47),X47)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
fof(c_0_43, plain, ![X67, X68]:(r1_xboole_0(X67,X68)|~r1_xboole_0(k3_xboole_0(X67,X68),X68)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t75_xboole_1])])])).
cnf(c_0_44, plain, (~r2_hidden(X1,k3_xboole_0(X2,k1_xboole_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_40]), c_0_41])).
cnf(c_0_45, plain, (X1=k1_xboole_0|r2_hidden(esk6_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_46, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(k3_xboole_0(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_47, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_44, c_0_45])).
fof(c_0_48, plain, ![X62, X63]:((~r1_xboole_0(X62,X63)|k4_xboole_0(X62,X63)=X62)&(k4_xboole_0(X62,X63)!=X62|r1_xboole_0(X62,X63))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
fof(c_0_49, negated_conjecture, ~(![X1, X2]:(k4_xboole_0(k1_tarski(X1),X2)=k1_xboole_0<=>r2_hidden(X1,X2))), inference(assume_negation,[status(cth)],[t68_zfmisc_1])).
fof(c_0_50, plain, ![X58, X59]:k5_xboole_0(X58,X59)=k5_xboole_0(X59,X58), inference(variable_rename,[status(thm)],[commutativity_k5_xboole_0])).
fof(c_0_51, plain, ![X8, X9]:k5_xboole_0(X8,X9)=k4_xboole_0(k2_xboole_0(X8,X9),k3_xboole_0(X8,X9)), inference(variable_rename,[status(thm)],[t101_xboole_1])).
cnf(c_0_52, plain, (r1_xboole_0(X1,k1_xboole_0)|~r1_xboole_0(k1_xboole_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_53, plain, (r1_xboole_0(X1,X2)|k4_xboole_0(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_54, negated_conjecture, ((k4_xboole_0(k1_tarski(esk1_0),esk2_0)!=k1_xboole_0|~r2_hidden(esk1_0,esk2_0))&(k4_xboole_0(k1_tarski(esk1_0),esk2_0)=k1_xboole_0|r2_hidden(esk1_0,esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_49])])])).
fof(c_0_55, plain, ![X7]:k1_enumset1(X7,X7,X7)=k1_tarski(X7), inference(variable_rename,[status(thm)],[t76_enumset1])).
fof(c_0_56, plain, ![X55, X56, X57]:k2_enumset1(X55,X55,X56,X57)=k1_enumset1(X55,X56,X57), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_57, plain, ![X69, X70, X71, X72]:k2_enumset1(X69,X70,X71,X72)=k2_xboole_0(k2_tarski(X69,X70),k2_tarski(X71,X72)), inference(variable_rename,[status(thm)],[t45_enumset1])).
fof(c_0_58, plain, ![X10, X11, X12, X13, X14, X15]:(((~r2_hidden(X12,X11)|X12=X10|X11!=k1_tarski(X10))&(X13!=X10|r2_hidden(X13,X11)|X11!=k1_tarski(X10)))&((~r2_hidden(esk3_2(X14,X15),X15)|esk3_2(X14,X15)!=X14|X15=k1_tarski(X14))&(r2_hidden(esk3_2(X14,X15),X15)|esk3_2(X14,X15)=X14|X15=k1_tarski(X14)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_59, plain, ![X20, X21]:(~r1_xboole_0(X20,X21)|r1_xboole_0(X21,X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
cnf(c_0_60, plain, (k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_61, plain, (k5_xboole_0(X1,X2)=k4_xboole_0(k2_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_62, plain, (k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_63, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_30])])).
cnf(c_0_64, negated_conjecture, (k4_xboole_0(k1_tarski(esk1_0),esk2_0)=k1_xboole_0|r2_hidden(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_65, plain, (k1_enumset1(X1,X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_66, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_67, plain, (k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X4))), inference(split_conjunct,[status(thm)],[c_0_57])).
fof(c_0_68, plain, ![X73]:k2_xboole_0(X73,X73)=X73, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k2_xboole_0])])).
cnf(c_0_69, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_58])).
fof(c_0_70, plain, ![X17, X18, X19]:k2_xboole_0(k2_tarski(X17,X18),X19)!=k1_xboole_0, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_zfmisc_1])])).
fof(c_0_71, plain, ![X74, X75, X76]:k3_xboole_0(X74,k4_xboole_0(X75,X76))=k4_xboole_0(k3_xboole_0(X74,X75),k3_xboole_0(X74,X76)), inference(variable_rename,[status(thm)],[t50_xboole_1])).
fof(c_0_72, plain, ![X49, X50, X51]:k3_xboole_0(X49,k4_xboole_0(X50,X51))=k4_xboole_0(k3_xboole_0(X49,X50),X51), inference(variable_rename,[status(thm)],[t49_xboole_1])).
fof(c_0_73, plain, ![X45, X46]:(~r1_xboole_0(X45,X46)|k4_xboole_0(k2_xboole_0(X45,X46),X46)=X45), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t88_xboole_1])])).
cnf(c_0_74, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_75, plain, (k4_xboole_0(k2_xboole_0(X1,X2),k3_xboole_0(X1,X2))=k4_xboole_0(k2_xboole_0(X2,X1),k3_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_61]), c_0_61])).
cnf(c_0_76, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_36, c_0_47])).
cnf(c_0_77, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_78, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_79, negated_conjecture, (k4_xboole_0(k2_xboole_0(k2_tarski(esk1_0,esk1_0),k2_tarski(esk1_0,esk1_0)),esk2_0)=k1_xboole_0|r2_hidden(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_65]), c_0_66]), c_0_67])).
cnf(c_0_80, plain, (k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_81, plain, (X1=X3|X2!=k2_xboole_0(k2_tarski(X3,X3),k2_tarski(X3,X3))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_65]), c_0_66]), c_0_67])).
cnf(c_0_82, plain, (k2_xboole_0(k2_tarski(X1,X2),X3)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_83, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_84, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_85, plain, (k4_xboole_0(k2_xboole_0(X1,X2),X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_86, plain, (r1_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_74, c_0_63])).
cnf(c_0_87, plain, (k2_xboole_0(k1_xboole_0,X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_38]), c_0_47]), c_0_77]), c_0_77])).
fof(c_0_88, plain, ![X64, X65, X66]:k2_xboole_0(X64,k3_xboole_0(X65,X66))=k3_xboole_0(k2_xboole_0(X64,X65),k2_xboole_0(X64,X66)), inference(variable_rename,[status(thm)],[t24_xboole_1])).
cnf(c_0_89, negated_conjecture, (k4_xboole_0(k1_tarski(esk1_0),esk2_0)!=k1_xboole_0|~r2_hidden(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_90, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_78])).
cnf(c_0_91, negated_conjecture, (k4_xboole_0(k2_tarski(esk1_0,esk1_0),esk2_0)=k1_xboole_0|r2_hidden(esk1_0,esk2_0)), inference(rw,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_92, plain, (X1=X2|~r2_hidden(X1,k2_tarski(X2,X2))), inference(er,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_80])])).
cnf(c_0_93, plain, (k2_tarski(X1,X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_82, c_0_38])).
cnf(c_0_94, plain, (k4_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3))=k4_xboole_0(k3_xboole_0(X1,X2),X3)), inference(rw,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_95, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_87])).
cnf(c_0_96, plain, (k2_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(k2_xboole_0(X1,X2),k2_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_97, negated_conjecture, (k4_xboole_0(k2_xboole_0(k2_tarski(esk1_0,esk1_0),k2_tarski(esk1_0,esk1_0)),esk2_0)!=k1_xboole_0|~r2_hidden(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_65]), c_0_66]), c_0_67])).
cnf(c_0_98, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|r2_hidden(X1,esk2_0)|~r2_hidden(X1,k2_tarski(esk1_0,esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_41])).
cnf(c_0_99, plain, (esk6_1(k2_tarski(X1,X1))=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_45]), c_0_93])).
cnf(c_0_100, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X2)=k1_xboole_0), inference(spm,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_101, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_38]), c_0_47]), c_0_38]), c_0_36])).
fof(c_0_102, plain, ![X42, X43, X44]:(~r2_hidden(X42,X43)|~r2_hidden(X44,X43)|k2_xboole_0(k2_tarski(X42,X44),X43)=X43), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t48_zfmisc_1])])).
cnf(c_0_103, negated_conjecture, (k4_xboole_0(k2_tarski(esk1_0,esk1_0),esk2_0)!=k1_xboole_0|~r2_hidden(esk1_0,esk2_0)), inference(rw,[status(thm)],[c_0_97, c_0_80])).
cnf(c_0_104, negated_conjecture, (r2_hidden(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_45]), c_0_99])]), c_0_93])).
cnf(c_0_105, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_106, plain, (k2_xboole_0(k2_tarski(X1,X3),X2)=X2|~r2_hidden(X1,X2)|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_107, negated_conjecture, (k4_xboole_0(k2_tarski(esk1_0,esk1_0),esk2_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_104])])).
cnf(c_0_108, plain, (k4_xboole_0(k2_tarski(X1,X2),X3)=k1_xboole_0|~r2_hidden(X2,X3)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_109, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_104])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 110
# Proof object clause steps            : 61
# Proof object formula steps           : 49
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 24
# Proof object generating inferences   : 22
# Proof object simplifying inferences  : 41
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 24
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 35
# Processed clauses                    : 1183
# ...of these trivial                  : 63
# ...subsumed                          : 863
# ...remaining for further processing  : 257
# Other redundant clauses eliminated   : 119
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 29
# Generated clauses                    : 7821
# ...of the previous two non-trivial   : 6099
# Contextual simplify-reflections      : 3
# Paramodulations                      : 7693
# Factorizations                       : 10
# Equation resolutions                 : 119
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 180
#    Positive orientable unit clauses  : 74
#    Positive unorientable unit clauses: 6
#    Negative unit clauses             : 11
#    Non-unit-clauses                  : 89
# Current number of unprocessed clauses: 4864
# ...number of literals in the above   : 9963
# Current number of archived formulas  : 0
# Current number of archived clauses   : 73
# Clause-clause subsumption calls (NU) : 4206
# Rec. Clause-clause subsumption calls : 3534
# Non-unit clause-clause subsumptions  : 230
# Unit Clause-clause subsumption calls : 396
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 160
# BW rewrite match successes           : 38
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 77831

# -------------------------------------------------
# User time                : 0.052 s
# System time              : 0.007 s
# Total time               : 0.059 s
# Maximum resident set size: 28684 pages
