# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t158_relat_1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>![X4]:(v1_relat_1(X4)=>((r1_tarski(X3,X4)&r1_tarski(X1,X2))=>r1_tarski(k9_relat_1(X3,X1),k9_relat_1(X4,X2))))), file('data/MPTP/proofs/1e65482dce7b4366a4e639d264cfcaa5.E_input', t158_relat_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/1e65482dce7b4366a4e639d264cfcaa5.E_input', d3_tarski)).
fof(d3_xboole_0, axiom, ![X1, X2, X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2)))), file('data/MPTP/proofs/1e65482dce7b4366a4e639d264cfcaa5.E_input', d3_xboole_0)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/1e65482dce7b4366a4e639d264cfcaa5.E_input', t1_xboole_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/1e65482dce7b4366a4e639d264cfcaa5.E_input', t7_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/1e65482dce7b4366a4e639d264cfcaa5.E_input', commutativity_k2_xboole_0)).
fof(t153_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>k9_relat_1(X3,k2_xboole_0(X1,X2))=k2_xboole_0(k9_relat_1(X3,X1),k9_relat_1(X3,X2))), file('data/MPTP/proofs/1e65482dce7b4366a4e639d264cfcaa5.E_input', t153_relat_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/1e65482dce7b4366a4e639d264cfcaa5.E_input', d10_xboole_0)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/1e65482dce7b4366a4e639d264cfcaa5.E_input', cc2_relat_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/1e65482dce7b4366a4e639d264cfcaa5.E_input', t3_subset)).
fof(t157_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(r1_tarski(X2,X3)=>r1_tarski(k9_relat_1(X2,X1),k9_relat_1(X3,X1))))), file('data/MPTP/proofs/1e65482dce7b4366a4e639d264cfcaa5.E_input', t157_relat_1)).
fof(c_0_11, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>![X4]:(v1_relat_1(X4)=>((r1_tarski(X3,X4)&r1_tarski(X1,X2))=>r1_tarski(k9_relat_1(X3,X1),k9_relat_1(X4,X2)))))), inference(assume_negation,[status(cth)],[t158_relat_1])).
fof(c_0_12, plain, ![X9, X10, X11, X12, X13]:((~r1_tarski(X9,X10)|(~r2_hidden(X11,X9)|r2_hidden(X11,X10)))&((r2_hidden(esk5_2(X12,X13),X12)|r1_tarski(X12,X13))&(~r2_hidden(esk5_2(X12,X13),X13)|r1_tarski(X12,X13)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_13, negated_conjecture, (v1_relat_1(esk3_0)&(v1_relat_1(esk4_0)&((r1_tarski(esk3_0,esk4_0)&r1_tarski(esk1_0,esk2_0))&~r1_tarski(k9_relat_1(esk3_0,esk1_0),k9_relat_1(esk4_0,esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
fof(c_0_14, plain, ![X15, X16, X17, X18, X19, X20, X21, X22]:(((~r2_hidden(X18,X17)|(r2_hidden(X18,X15)|r2_hidden(X18,X16))|X17!=k2_xboole_0(X15,X16))&((~r2_hidden(X19,X15)|r2_hidden(X19,X17)|X17!=k2_xboole_0(X15,X16))&(~r2_hidden(X19,X16)|r2_hidden(X19,X17)|X17!=k2_xboole_0(X15,X16))))&(((~r2_hidden(esk6_3(X20,X21,X22),X20)|~r2_hidden(esk6_3(X20,X21,X22),X22)|X22=k2_xboole_0(X20,X21))&(~r2_hidden(esk6_3(X20,X21,X22),X21)|~r2_hidden(esk6_3(X20,X21,X22),X22)|X22=k2_xboole_0(X20,X21)))&(r2_hidden(esk6_3(X20,X21,X22),X22)|(r2_hidden(esk6_3(X20,X21,X22),X20)|r2_hidden(esk6_3(X20,X21,X22),X21))|X22=k2_xboole_0(X20,X21)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])).
cnf(c_0_15, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X2!=k2_xboole_0(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_20, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k2_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_21, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_22, plain, ![X24, X25, X26]:(~r1_tarski(X24,X25)|~r1_tarski(X25,X26)|r1_tarski(X24,X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_23, plain, ![X27, X28]:r1_tarski(X27,k2_xboole_0(X27,X28)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_24, negated_conjecture, (r1_tarski(X1,esk2_0)|~r2_hidden(esk5_2(X1,esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25, plain, (r2_hidden(esk5_2(k2_xboole_0(X1,X2),X3),X2)|r2_hidden(esk5_2(k2_xboole_0(X1,X2),X3),X1)|r1_tarski(k2_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
fof(c_0_26, plain, ![X29, X30]:k2_xboole_0(X29,X30)=k2_xboole_0(X30,X29), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_27, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_29, plain, ![X31, X32, X33]:(~v1_relat_1(X33)|k9_relat_1(X33,k2_xboole_0(X31,X32))=k2_xboole_0(k9_relat_1(X33,X31),k9_relat_1(X33,X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t153_relat_1])])).
fof(c_0_30, plain, ![X34, X35]:(((r1_tarski(X34,X35)|X34!=X35)&(r1_tarski(X35,X34)|X34!=X35))&(~r1_tarski(X34,X35)|~r1_tarski(X35,X34)|X34=X35)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_31, negated_conjecture, (r2_hidden(esk5_2(k2_xboole_0(esk1_0,X1),esk2_0),X1)|r1_tarski(k2_xboole_0(esk1_0,X1),esk2_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_32, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_33, plain, ![X36, X37]:(~v1_relat_1(X36)|(~m1_subset_1(X37,k1_zfmisc_1(X36))|v1_relat_1(X37))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_34, plain, ![X38, X39]:((~m1_subset_1(X38,k1_zfmisc_1(X39))|r1_tarski(X38,X39))&(~r1_tarski(X38,X39)|m1_subset_1(X38,k1_zfmisc_1(X39)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_35, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_36, plain, (k9_relat_1(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k9_relat_1(X1,X2),k9_relat_1(X1,X3))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, negated_conjecture, (r1_tarski(k2_xboole_0(esk1_0,esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_18, c_0_31])).
cnf(c_0_39, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_28, c_0_32])).
fof(c_0_40, plain, ![X40, X41, X42]:(~v1_relat_1(X41)|(~v1_relat_1(X42)|(~r1_tarski(X41,X42)|r1_tarski(k9_relat_1(X41,X40),k9_relat_1(X42,X40))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t157_relat_1])])])).
cnf(c_0_41, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43, plain, (r1_tarski(X1,k9_relat_1(X2,k2_xboole_0(X3,X4)))|~r1_tarski(X1,k9_relat_1(X2,X3))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_44, negated_conjecture, (k2_xboole_0(esk1_0,esk2_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])])).
cnf(c_0_45, plain, (r1_tarski(k9_relat_1(X1,X3),k9_relat_1(X2,X3))|~v1_relat_1(X1)|~v1_relat_1(X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_46, plain, (v1_relat_1(X1)|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_47, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_48, negated_conjecture, (~r1_tarski(k9_relat_1(esk3_0,esk1_0),k9_relat_1(esk4_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_49, negated_conjecture, (r1_tarski(X1,k9_relat_1(X2,esk2_0))|~r1_tarski(X1,k9_relat_1(X2,esk1_0))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_50, negated_conjecture, (v1_relat_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_51, plain, (r1_tarski(X1,k9_relat_1(X2,X3))|~r1_tarski(X1,k9_relat_1(X4,X3))|~r1_tarski(X4,X2)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_45]), c_0_46])).
cnf(c_0_52, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_47])).
cnf(c_0_53, negated_conjecture, (~r1_tarski(k9_relat_1(esk3_0,esk1_0),k9_relat_1(esk4_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])])).
cnf(c_0_54, plain, (r1_tarski(k9_relat_1(X1,X2),k9_relat_1(X3,X2))|~r1_tarski(X1,X3)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_55, negated_conjecture, (r1_tarski(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_56, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55]), c_0_50])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 57
# Proof object clause steps            : 34
# Proof object formula steps           : 23
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 10
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 25
# Processed clauses                    : 13940
# ...of these trivial                  : 154
# ...subsumed                          : 11016
# ...remaining for further processing  : 2770
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 118
# Backward-rewritten                   : 13
# Generated clauses                    : 102228
# ...of the previous two non-trivial   : 96082
# Contextual simplify-reflections      : 341
# Paramodulations                      : 101669
# Factorizations                       : 554
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2610
#    Positive orientable unit clauses  : 128
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 2478
# Current number of unprocessed clauses: 81939
# ...number of literals in the above   : 346911
# Current number of archived formulas  : 0
# Current number of archived clauses   : 155
# Clause-clause subsumption calls (NU) : 1093511
# Rec. Clause-clause subsumption calls : 468335
# Non-unit clause-clause subsumptions  : 10966
# Unit Clause-clause subsumption calls : 9032
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3752
# BW rewrite match successes           : 13
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2069555

# -------------------------------------------------
# User time                : 1.398 s
# System time              : 0.017 s
# Total time               : 1.415 s
# Maximum resident set size: 68692 pages
