# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', t100_xboole_1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', t12_setfam_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', t3_boole)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', t2_boole)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', d3_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', dt_l1_pre_topc)).
fof(d5_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)), file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', d5_subset_1)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', d4_subset_1)).
fof(t80_tops_1, conjecture, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v1_tops_1(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>~((~(X3=k1_xboole_0)&(v3_pre_topc(X3,X1)&r1_xboole_0(X2,X3)))))))), file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', t80_tops_1)).
fof(d3_tops_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v1_tops_1(X2,X1)<=>k2_pre_topc(X1,X2)=k2_struct_0(X1)))), file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', d3_tops_1)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', t83_xboole_1)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', t4_subset_1)).
fof(d13_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k2_pre_topc(X1,X2)<=>![X4]:(r2_hidden(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))=>~((v3_pre_topc(X5,X1)&(r2_hidden(X4,X5)&r1_xboole_0(X2,X5))))))))))), file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', d13_pre_topc)).
fof(involutiveness_k3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,k3_subset_1(X1,X2))=X2), file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', involutiveness_k3_subset_1)).
fof(t29_tops_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v4_pre_topc(X2,X1)<=>v3_pre_topc(k3_subset_1(u1_struct_0(X1),X2),X1)))), file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', t29_tops_1)).
fof(t52_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v4_pre_topc(X2,X1)=>k2_pre_topc(X1,X2)=X2)&((v2_pre_topc(X1)&k2_pre_topc(X1,X2)=X2)=>v4_pre_topc(X2,X1))))), file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', t52_pre_topc)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', d7_xboole_0)).
fof(t54_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k2_pre_topc(X1,X2))<=>(~(v2_struct_0(X1))&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>~((v3_pre_topc(X4,X1)&(r2_hidden(X3,X4)&r1_xboole_0(X2,X4)))))))))), file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', t54_pre_topc)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', t4_subset)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', l3_subset_1)).
fof(fc10_tops_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>v3_pre_topc(k2_struct_0(X1),X1)), file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', fc10_tops_1)).
fof(t4_boole, axiom, ![X1]:k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0, file('data/MPTP/proofs/1e9a2a5b99af48dda2da38825fa4ebf7.E_input', t4_boole)).
fof(c_0_23, plain, ![X50, X51]:k4_xboole_0(X50,X51)=k5_xboole_0(X50,k3_xboole_0(X50,X51)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_24, plain, ![X14, X15]:k1_setfam_1(k2_tarski(X14,X15))=k3_xboole_0(X14,X15), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_25, plain, ![X44]:k4_xboole_0(X44,k1_xboole_0)=X44, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_26, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_28, plain, ![X10]:k3_xboole_0(X10,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_29, plain, ![X13]:(~l1_struct_0(X13)|k2_struct_0(X13)=u1_struct_0(X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
fof(c_0_30, plain, ![X56]:(~l1_pre_topc(X56)|l1_struct_0(X56)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_31, plain, ![X29, X30]:(~m1_subset_1(X30,k1_zfmisc_1(X29))|k3_subset_1(X29,X30)=k4_xboole_0(X29,X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
cnf(c_0_32, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X2)))), inference(rw,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_34, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_35, plain, ![X36]:m1_subset_1(k2_subset_1(X36),k1_zfmisc_1(X36)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_36, plain, ![X16]:k2_subset_1(X16)=X16, inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_37, negated_conjecture, ~(![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v1_tops_1(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>~((~(X3=k1_xboole_0)&(v3_pre_topc(X3,X1)&r1_xboole_0(X2,X3))))))))), inference(assume_negation,[status(cth)],[t80_tops_1])).
fof(c_0_38, plain, ![X42, X43]:((~v1_tops_1(X43,X42)|k2_pre_topc(X42,X43)=k2_struct_0(X42)|~m1_subset_1(X43,k1_zfmisc_1(u1_struct_0(X42)))|~l1_pre_topc(X42))&(k2_pre_topc(X42,X43)!=k2_struct_0(X42)|v1_tops_1(X43,X42)|~m1_subset_1(X43,k1_zfmisc_1(u1_struct_0(X42)))|~l1_pre_topc(X42))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tops_1])])])])).
cnf(c_0_39, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_40, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_41, plain, ![X11, X12]:((~r1_xboole_0(X11,X12)|k4_xboole_0(X11,X12)=X11)&(k4_xboole_0(X11,X12)!=X11|r1_xboole_0(X11,X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
cnf(c_0_42, plain, (k3_subset_1(X2,X1)=k4_xboole_0(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_43, plain, (k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,k1_xboole_0)))=X1), inference(rw,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_44, plain, (k1_setfam_1(k2_tarski(X1,k1_xboole_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_34, c_0_27])).
fof(c_0_45, plain, ![X33]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X33)), inference(variable_rename,[status(thm)],[t4_subset_1])).
fof(c_0_46, plain, ![X20, X21, X22, X23, X24, X28]:(((~r2_hidden(X23,X22)|(~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X20)))|(~v3_pre_topc(X24,X20)|(~r2_hidden(X23,X24)|~r1_xboole_0(X21,X24))))|~r2_hidden(X23,u1_struct_0(X20))|X22!=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))&((m1_subset_1(esk4_4(X20,X21,X22,X23),k1_zfmisc_1(u1_struct_0(X20)))|r2_hidden(X23,X22)|~r2_hidden(X23,u1_struct_0(X20))|X22!=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))&((v3_pre_topc(esk4_4(X20,X21,X22,X23),X20)|r2_hidden(X23,X22)|~r2_hidden(X23,u1_struct_0(X20))|X22!=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))&((r2_hidden(X23,esk4_4(X20,X21,X22,X23))|r2_hidden(X23,X22)|~r2_hidden(X23,u1_struct_0(X20))|X22!=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))&(r1_xboole_0(X21,esk4_4(X20,X21,X22,X23))|r2_hidden(X23,X22)|~r2_hidden(X23,u1_struct_0(X20))|X22!=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))))))&((r2_hidden(esk5_3(X20,X21,X22),u1_struct_0(X20))|X22=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))&(((m1_subset_1(esk6_3(X20,X21,X22),k1_zfmisc_1(u1_struct_0(X20)))|~r2_hidden(esk5_3(X20,X21,X22),X22)|X22=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))&((v3_pre_topc(esk6_3(X20,X21,X22),X20)|~r2_hidden(esk5_3(X20,X21,X22),X22)|X22=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))&((r2_hidden(esk5_3(X20,X21,X22),esk6_3(X20,X21,X22))|~r2_hidden(esk5_3(X20,X21,X22),X22)|X22=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))&(r1_xboole_0(X21,esk6_3(X20,X21,X22))|~r2_hidden(esk5_3(X20,X21,X22),X22)|X22=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20)))))&(r2_hidden(esk5_3(X20,X21,X22),X22)|(~m1_subset_1(X28,k1_zfmisc_1(u1_struct_0(X20)))|(~v3_pre_topc(X28,X20)|(~r2_hidden(esk5_3(X20,X21,X22),X28)|~r1_xboole_0(X21,X28))))|X22=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d13_pre_topc])])])])])).
cnf(c_0_47, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_48, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_49, negated_conjecture, ![X9]:((v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(((m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_tops_1(esk2_0,esk1_0))&((esk3_0!=k1_xboole_0|~v1_tops_1(esk2_0,esk1_0))&((v3_pre_topc(esk3_0,esk1_0)|~v1_tops_1(esk2_0,esk1_0))&(r1_xboole_0(esk2_0,esk3_0)|~v1_tops_1(esk2_0,esk1_0)))))&(v1_tops_1(esk2_0,esk1_0)|(~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(esk1_0)))|(X9=k1_xboole_0|(~v3_pre_topc(X9,esk1_0)|~r1_xboole_0(esk2_0,X9)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_37])])])])])])).
cnf(c_0_50, plain, (v1_tops_1(X2,X1)|k2_pre_topc(X1,X2)!=k2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_51, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_52, plain, (k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_53, plain, ![X52, X53]:(~m1_subset_1(X53,k1_zfmisc_1(X52))|k3_subset_1(X52,k3_subset_1(X52,X53))=X53), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k3_subset_1])])).
cnf(c_0_54, plain, (k3_subset_1(X2,X1)=k5_xboole_0(X2,k1_setfam_1(k2_tarski(X2,X1)))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_42, c_0_33])).
cnf(c_0_55, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_56, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_57, plain, (v3_pre_topc(esk6_3(X1,X2,X3),X1)|X3=k2_pre_topc(X1,X2)|~r2_hidden(esk5_3(X1,X2,X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_58, plain, (r2_hidden(esk5_3(X1,X2,X3),u1_struct_0(X1))|X3=k2_pre_topc(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_59, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_60, negated_conjecture, (v3_pre_topc(esk3_0,esk1_0)|~v1_tops_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_61, plain, (v1_tops_1(X1,X2)|k2_pre_topc(X2,X1)!=u1_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_62, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_63, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_64, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_tops_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_65, negated_conjecture, (r1_xboole_0(esk2_0,esk3_0)|~v1_tops_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_66, negated_conjecture, (esk3_0!=k1_xboole_0|~v1_tops_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_67, plain, (k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X2)))=X1|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_52, c_0_33])).
cnf(c_0_68, plain, (k3_subset_1(X2,k3_subset_1(X2,X1))=X1|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_69, plain, (k3_subset_1(X1,k1_xboole_0)=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_44]), c_0_55]), c_0_56])])).
cnf(c_0_70, negated_conjecture, (v1_tops_1(esk2_0,esk1_0)|X1=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v3_pre_topc(X1,esk1_0)|~r1_xboole_0(esk2_0,X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_71, plain, (k2_pre_topc(X1,X2)=u1_struct_0(X1)|v3_pre_topc(esk6_3(X1,X2,u1_struct_0(X1)),X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59])])).
cnf(c_0_72, plain, (r1_xboole_0(X1,esk6_3(X2,X1,X3))|X3=k2_pre_topc(X2,X1)|~r2_hidden(esk5_3(X2,X1,X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_73, negated_conjecture, (v3_pre_topc(esk3_0,esk1_0)|k2_pre_topc(esk1_0,esk2_0)!=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62]), c_0_63])])).
cnf(c_0_74, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|k2_pre_topc(esk1_0,esk2_0)!=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_61]), c_0_62]), c_0_63])])).
cnf(c_0_75, negated_conjecture, (r1_xboole_0(esk2_0,esk3_0)|k2_pre_topc(esk1_0,esk2_0)!=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_61]), c_0_62]), c_0_63])])).
cnf(c_0_76, negated_conjecture, (k2_pre_topc(esk1_0,esk2_0)!=u1_struct_0(esk1_0)|esk3_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_61]), c_0_62]), c_0_63])])).
cnf(c_0_77, plain, (k3_subset_1(X1,X2)=X1|~r1_xboole_0(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_67, c_0_54])).
cnf(c_0_78, plain, (k3_subset_1(X1,X1)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_56])])).
fof(c_0_79, plain, ![X31, X32]:((~v4_pre_topc(X32,X31)|v3_pre_topc(k3_subset_1(u1_struct_0(X31),X32),X31)|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|~l1_pre_topc(X31))&(~v3_pre_topc(k3_subset_1(u1_struct_0(X31),X32),X31)|v4_pre_topc(X32,X31)|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|~l1_pre_topc(X31))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t29_tops_1])])])])).
cnf(c_0_80, negated_conjecture, (esk6_3(esk1_0,X1,u1_struct_0(esk1_0))=k1_xboole_0|k2_pre_topc(esk1_0,X1)=u1_struct_0(esk1_0)|v1_tops_1(esk2_0,esk1_0)|~r1_xboole_0(esk2_0,esk6_3(esk1_0,X1,u1_struct_0(esk1_0)))|~m1_subset_1(esk6_3(esk1_0,X1,u1_struct_0(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_63])])).
cnf(c_0_81, plain, (k2_pre_topc(X1,X2)=u1_struct_0(X1)|r1_xboole_0(X2,esk6_3(X1,X2,u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_58]), c_0_59])])).
cnf(c_0_82, negated_conjecture, (v1_tops_1(esk2_0,esk1_0)|k2_pre_topc(esk1_0,esk2_0)!=u1_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_73]), c_0_74]), c_0_75]), c_0_76])).
cnf(c_0_83, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|X3=k2_pre_topc(X1,X2)|~r2_hidden(esk5_3(X1,X2,X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_84, plain, (k1_xboole_0=X1|~r1_xboole_0(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_77]), c_0_78])).
fof(c_0_85, plain, ![X34, X35]:((~v4_pre_topc(X35,X34)|k2_pre_topc(X34,X35)=X35|~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X34)))|~l1_pre_topc(X34))&(~v2_pre_topc(X34)|k2_pre_topc(X34,X35)!=X35|v4_pre_topc(X35,X34)|~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X34)))|~l1_pre_topc(X34))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t52_pre_topc])])])])).
cnf(c_0_86, plain, (v4_pre_topc(X2,X1)|~v3_pre_topc(k3_subset_1(u1_struct_0(X1),X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_87, negated_conjecture, (esk6_3(esk1_0,esk2_0,u1_struct_0(esk1_0))=k1_xboole_0|v1_tops_1(esk2_0,esk1_0)|~m1_subset_1(esk6_3(esk1_0,esk2_0,u1_struct_0(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_62]), c_0_63])]), c_0_82])).
cnf(c_0_88, plain, (k2_pre_topc(X1,X2)=u1_struct_0(X1)|m1_subset_1(esk6_3(X1,X2,u1_struct_0(X1)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_58]), c_0_59])])).
cnf(c_0_89, plain, (k2_pre_topc(X2,X1)=k2_struct_0(X2)|~v1_tops_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_90, plain, (esk6_3(X1,X2,u1_struct_0(X1))=k1_xboole_0|k2_pre_topc(X1,X2)=u1_struct_0(X1)|~m1_subset_1(esk6_3(X1,X2,u1_struct_0(X1)),k1_zfmisc_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_84, c_0_81])).
cnf(c_0_91, plain, (k2_pre_topc(X2,X1)=X1|~v4_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_92, plain, (v4_pre_topc(u1_struct_0(X1),X1)|~v3_pre_topc(k1_xboole_0,X1)|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_78]), c_0_59])])).
cnf(c_0_93, negated_conjecture, (esk6_3(esk1_0,esk2_0,u1_struct_0(esk1_0))=k1_xboole_0|v1_tops_1(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_62]), c_0_63])]), c_0_82])).
fof(c_0_94, plain, ![X45, X46]:((~r1_xboole_0(X45,X46)|k3_xboole_0(X45,X46)=k1_xboole_0)&(k3_xboole_0(X45,X46)!=k1_xboole_0|r1_xboole_0(X45,X46))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_95, negated_conjecture, (k2_struct_0(esk1_0)=k2_pre_topc(esk1_0,esk2_0)|k2_pre_topc(esk1_0,esk2_0)!=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_82]), c_0_62]), c_0_63])])).
cnf(c_0_96, plain, (esk6_3(X1,u1_struct_0(X1),u1_struct_0(X1))=k1_xboole_0|k2_pre_topc(X1,u1_struct_0(X1))=u1_struct_0(X1)|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_88]), c_0_59])])).
cnf(c_0_97, plain, (k2_pre_topc(X1,u1_struct_0(X1))=u1_struct_0(X1)|~v3_pre_topc(k1_xboole_0,X1)|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_59])])).
cnf(c_0_98, negated_conjecture, (v3_pre_topc(k1_xboole_0,esk1_0)|v1_tops_1(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_93]), c_0_62]), c_0_63])]), c_0_82])).
cnf(c_0_99, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_100, negated_conjecture, (v1_tops_1(X1,esk1_0)|k2_pre_topc(esk1_0,X1)!=k2_pre_topc(esk1_0,esk2_0)|k2_pre_topc(esk1_0,esk2_0)!=u1_struct_0(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_95]), c_0_63])])).
cnf(c_0_101, plain, (k2_pre_topc(X1,u1_struct_0(X1))=u1_struct_0(X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_96]), c_0_59])]), c_0_97])).
cnf(c_0_102, plain, (v3_pre_topc(k3_subset_1(u1_struct_0(X2),X1),X2)|~v4_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_103, plain, (v4_pre_topc(k3_subset_1(u1_struct_0(X1),X2),X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(k3_subset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_86, c_0_68])).
cnf(c_0_104, negated_conjecture, (k2_struct_0(esk1_0)=k2_pre_topc(esk1_0,esk2_0)|v3_pre_topc(k1_xboole_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_98]), c_0_62]), c_0_63])])).
cnf(c_0_105, plain, (r1_xboole_0(X1,X2)|k1_setfam_1(k2_tarski(X1,X2))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_99, c_0_27])).
cnf(c_0_106, negated_conjecture, (k2_pre_topc(esk1_0,X1)=k2_struct_0(esk1_0)|k2_pre_topc(esk1_0,X1)!=k2_pre_topc(esk1_0,esk2_0)|k2_pre_topc(esk1_0,esk2_0)!=u1_struct_0(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_100]), c_0_63])])).
cnf(c_0_107, plain, (v1_tops_1(u1_struct_0(X1),X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_101]), c_0_59])]), c_0_51])).
cnf(c_0_108, plain, (v3_pre_topc(X1,X2)|~v4_pre_topc(k3_subset_1(u1_struct_0(X2),X1),X2)|~m1_subset_1(k3_subset_1(u1_struct_0(X2),X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_102, c_0_68])).
cnf(c_0_109, plain, (v4_pre_topc(X2,X1)|~v2_pre_topc(X1)|k2_pre_topc(X1,X2)!=X2|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_110, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_111, plain, (v4_pre_topc(u1_struct_0(X1),X1)|~r1_xboole_0(u1_struct_0(X1),X2)|~v3_pre_topc(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_77]), c_0_59])])).
cnf(c_0_112, negated_conjecture, (k2_pre_topc(esk1_0,esk2_0)=u1_struct_0(esk1_0)|v3_pre_topc(k1_xboole_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_104]), c_0_63])])).
cnf(c_0_113, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_105, c_0_44])).
cnf(c_0_114, negated_conjecture, (k2_struct_0(esk1_0)=u1_struct_0(esk1_0)|k2_pre_topc(esk1_0,esk2_0)!=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_101]), c_0_59]), c_0_63])])).
cnf(c_0_115, plain, (k2_pre_topc(X1,u1_struct_0(X1))=k2_struct_0(X1)|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_107]), c_0_59])])).
fof(c_0_116, plain, ![X37, X38, X39, X40]:(((~v2_struct_0(X37)|~r2_hidden(X39,k2_pre_topc(X37,X38))|~m1_subset_1(X39,u1_struct_0(X37))|~m1_subset_1(X38,k1_zfmisc_1(u1_struct_0(X37)))|~l1_pre_topc(X37))&(~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X37)))|(~v3_pre_topc(X40,X37)|(~r2_hidden(X39,X40)|~r1_xboole_0(X38,X40)))|~r2_hidden(X39,k2_pre_topc(X37,X38))|~m1_subset_1(X39,u1_struct_0(X37))|~m1_subset_1(X38,k1_zfmisc_1(u1_struct_0(X37)))|~l1_pre_topc(X37)))&((m1_subset_1(esk7_3(X37,X38,X39),k1_zfmisc_1(u1_struct_0(X37)))|v2_struct_0(X37)|r2_hidden(X39,k2_pre_topc(X37,X38))|~m1_subset_1(X39,u1_struct_0(X37))|~m1_subset_1(X38,k1_zfmisc_1(u1_struct_0(X37)))|~l1_pre_topc(X37))&((v3_pre_topc(esk7_3(X37,X38,X39),X37)|v2_struct_0(X37)|r2_hidden(X39,k2_pre_topc(X37,X38))|~m1_subset_1(X39,u1_struct_0(X37))|~m1_subset_1(X38,k1_zfmisc_1(u1_struct_0(X37)))|~l1_pre_topc(X37))&((r2_hidden(X39,esk7_3(X37,X38,X39))|v2_struct_0(X37)|r2_hidden(X39,k2_pre_topc(X37,X38))|~m1_subset_1(X39,u1_struct_0(X37))|~m1_subset_1(X38,k1_zfmisc_1(u1_struct_0(X37)))|~l1_pre_topc(X37))&(r1_xboole_0(X38,esk7_3(X37,X38,X39))|v2_struct_0(X37)|r2_hidden(X39,k2_pre_topc(X37,X38))|~m1_subset_1(X39,u1_struct_0(X37))|~m1_subset_1(X38,k1_zfmisc_1(u1_struct_0(X37)))|~l1_pre_topc(X37)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t54_pre_topc])])])])])])).
fof(c_0_117, plain, ![X47, X48, X49]:(~r2_hidden(X47,X48)|~m1_subset_1(X48,k1_zfmisc_1(X49))|m1_subset_1(X47,X49)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_118, plain, (v3_pre_topc(X1,X2)|~v4_pre_topc(u1_struct_0(X2),X2)|~r1_xboole_0(u1_struct_0(X2),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_77]), c_0_59])])).
cnf(c_0_119, negated_conjecture, (v4_pre_topc(X1,esk1_0)|k2_pre_topc(esk1_0,X1)!=X1|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_63])])).
cnf(c_0_120, negated_conjecture, (k2_pre_topc(esk1_0,esk2_0)=u1_struct_0(esk1_0)|v4_pre_topc(u1_struct_0(esk1_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_113]), c_0_56]), c_0_63])])).
cnf(c_0_121, negated_conjecture, (k2_pre_topc(esk1_0,u1_struct_0(esk1_0))=u1_struct_0(esk1_0)|k2_pre_topc(esk1_0,esk2_0)!=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_63])])).
cnf(c_0_122, plain, (~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v3_pre_topc(X1,X2)|~r2_hidden(X3,X1)|~r1_xboole_0(X4,X1)|~r2_hidden(X3,k2_pre_topc(X2,X4))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_116])).
cnf(c_0_123, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_117])).
cnf(c_0_124, negated_conjecture, (v3_pre_topc(X1,esk1_0)|k2_pre_topc(esk1_0,u1_struct_0(esk1_0))!=u1_struct_0(esk1_0)|~r1_xboole_0(u1_struct_0(esk1_0),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_63]), c_0_59])])).
cnf(c_0_125, negated_conjecture, (k2_pre_topc(esk1_0,u1_struct_0(esk1_0))=u1_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_120]), c_0_59]), c_0_63])]), c_0_121])).
fof(c_0_126, plain, ![X17, X18, X19]:(~m1_subset_1(X18,k1_zfmisc_1(X17))|(~r2_hidden(X19,X18)|r2_hidden(X19,X17))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
cnf(c_0_127, plain, (~r2_hidden(X1,k2_pre_topc(X2,X3))|~r2_hidden(X1,X4)|~r1_xboole_0(X3,X4)|~v3_pre_topc(X4,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(csr,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_128, negated_conjecture, (v3_pre_topc(X1,esk1_0)|~r1_xboole_0(u1_struct_0(esk1_0),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_124, c_0_125])])).
cnf(c_0_129, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_126])).
cnf(c_0_130, negated_conjecture, (~r2_hidden(X1,X2)|~r1_xboole_0(u1_struct_0(esk1_0),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_125]), c_0_59]), c_0_63])]), c_0_128]), c_0_129])).
cnf(c_0_131, plain, (r2_hidden(esk5_3(X1,X2,X3),esk6_3(X1,X2,X3))|X3=k2_pre_topc(X1,X2)|~r2_hidden(esk5_3(X1,X2,X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_132, negated_conjecture, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_113]), c_0_56])])).
cnf(c_0_133, negated_conjecture, (v1_tops_1(esk2_0,esk1_0)|~r2_hidden(esk5_3(esk1_0,esk2_0,u1_struct_0(esk1_0)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_93]), c_0_59]), c_0_62]), c_0_63])]), c_0_132]), c_0_82])).
cnf(c_0_134, plain, (k2_pre_topc(X1,X2)=k2_struct_0(X1)|k2_pre_topc(X1,X2)!=u1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_89, c_0_61])).
fof(c_0_135, plain, ![X55]:(~v2_pre_topc(X55)|~l1_pre_topc(X55)|v3_pre_topc(k2_struct_0(X55),X55)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc10_tops_1])])).
cnf(c_0_136, negated_conjecture, (v1_tops_1(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_58]), c_0_59]), c_0_62]), c_0_63])]), c_0_82])).
cnf(c_0_137, negated_conjecture, (k2_struct_0(esk1_0)=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_125]), c_0_59]), c_0_63])])).
cnf(c_0_138, plain, (v3_pre_topc(k2_struct_0(X1),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_135])).
cnf(c_0_139, negated_conjecture, (k2_pre_topc(esk1_0,esk2_0)=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_136]), c_0_137]), c_0_62]), c_0_63])])).
cnf(c_0_140, plain, (r2_hidden(esk5_3(X1,X2,X3),X3)|X3=k2_pre_topc(X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v3_pre_topc(X4,X1)|~r2_hidden(esk5_3(X1,X2,X3),X4)|~r1_xboole_0(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_141, plain, (v3_pre_topc(u1_struct_0(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_138, c_0_51])).
fof(c_0_142, plain, ![X54]:k4_xboole_0(k1_xboole_0,X54)=k1_xboole_0, inference(variable_rename,[status(thm)],[t4_boole])).
cnf(c_0_143, negated_conjecture, (~r2_hidden(X1,X2)|~r1_xboole_0(esk2_0,X2)|~v3_pre_topc(X2,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_139]), c_0_62]), c_0_63])]), c_0_129])).
cnf(c_0_144, negated_conjecture, (v3_pre_topc(esk3_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_136])])).
cnf(c_0_145, negated_conjecture, (r1_xboole_0(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_136])])).
cnf(c_0_146, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_136])])).
cnf(c_0_147, plain, (X1=k2_pre_topc(X2,X3)|r2_hidden(esk5_3(X2,X3,X1),X1)|~r1_xboole_0(X3,u1_struct_0(X2))|~v3_pre_topc(u1_struct_0(X2),X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_58]), c_0_59])])).
cnf(c_0_148, negated_conjecture, (v3_pre_topc(u1_struct_0(esk1_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_110]), c_0_63])])).
cnf(c_0_149, plain, (r1_xboole_0(X1,X2)|k4_xboole_0(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_150, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_142])).
cnf(c_0_151, plain, (v4_pre_topc(k1_xboole_0,X1)|~v3_pre_topc(u1_struct_0(X1),X1)|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_69]), c_0_56])])).
cnf(c_0_152, negated_conjecture, (~r2_hidden(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144]), c_0_145]), c_0_146])])).
cnf(c_0_153, negated_conjecture, (X1=k2_pre_topc(esk1_0,X2)|r2_hidden(esk5_3(esk1_0,X2,X1),X1)|~r1_xboole_0(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_148]), c_0_63])])).
cnf(c_0_154, plain, (r1_xboole_0(X1,X2)|k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X2)))!=X1), inference(rw,[status(thm)],[c_0_149, c_0_33])).
cnf(c_0_155, plain, (k5_xboole_0(k1_xboole_0,k1_setfam_1(k2_tarski(k1_xboole_0,X1)))=k1_xboole_0), inference(rw,[status(thm)],[c_0_150, c_0_33])).
cnf(c_0_156, negated_conjecture, (v4_pre_topc(k1_xboole_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_148]), c_0_63])])).
cnf(c_0_157, negated_conjecture, (k2_pre_topc(esk1_0,X1)=esk3_0|~r1_xboole_0(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_153]), c_0_146])])).
cnf(c_0_158, plain, (r1_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_154, c_0_155])).
cnf(c_0_159, negated_conjecture, (k2_pre_topc(esk1_0,k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_156]), c_0_56]), c_0_63])])).
cnf(c_0_160, negated_conjecture, (esk3_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_136])])).
cnf(c_0_161, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_158]), c_0_159]), c_0_56])]), c_0_160]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 162
# Proof object clause steps            : 115
# Proof object formula steps           : 47
# Proof object conjectures             : 50
# Proof object clause conjectures      : 47
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 39
# Proof object initial formulas used   : 23
# Proof object generating inferences   : 60
# Proof object simplifying inferences  : 161
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 23
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 50
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 47
# Processed clauses                    : 1461
# ...of these trivial                  : 30
# ...subsumed                          : 839
# ...remaining for further processing  : 592
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 191
# Backward-rewritten                   : 80
# Generated clauses                    : 2686
# ...of the previous two non-trivial   : 2266
# Contextual simplify-reflections      : 106
# Paramodulations                      : 2671
# Factorizations                       : 0
# Equation resolutions                 : 15
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 270
#    Positive orientable unit clauses  : 24
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 240
# Current number of unprocessed clauses: 696
# ...number of literals in the above   : 5566
# Current number of archived formulas  : 0
# Current number of archived clauses   : 320
# Clause-clause subsumption calls (NU) : 59426
# Rec. Clause-clause subsumption calls : 8286
# Non-unit clause-clause subsumptions  : 933
# Unit Clause-clause subsumption calls : 828
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 18
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 81546

# -------------------------------------------------
# User time                : 0.076 s
# System time              : 0.007 s
# Total time               : 0.082 s
# Maximum resident set size: 28796 pages
