# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t106_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k4_xboole_0(X2,X3))=>(r1_tarski(X1,X2)&r1_xboole_0(X1,X3))), file('data/MPTP/proofs/1fd4e0c15f88414d8c964933c674d960.E_input', t106_xboole_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/1fd4e0c15f88414d8c964933c674d960.E_input', t36_xboole_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/1fd4e0c15f88414d8c964933c674d960.E_input', t3_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/1fd4e0c15f88414d8c964933c674d960.E_input', t28_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/1fd4e0c15f88414d8c964933c674d960.E_input', t48_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/1fd4e0c15f88414d8c964933c674d960.E_input', commutativity_k3_xboole_0)).
fof(t149_zfmisc_1, conjecture, ![X1, X2, X3, X4]:((r1_tarski(k3_xboole_0(X1,X2),k1_tarski(X4))&(r2_hidden(X4,X3)&r1_xboole_0(X3,X2)))=>r1_xboole_0(k2_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/1fd4e0c15f88414d8c964933c674d960.E_input', t149_zfmisc_1)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/1fd4e0c15f88414d8c964933c674d960.E_input', t83_xboole_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/1fd4e0c15f88414d8c964933c674d960.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/1fd4e0c15f88414d8c964933c674d960.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/1fd4e0c15f88414d8c964933c674d960.E_input', t71_enumset1)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/1fd4e0c15f88414d8c964933c674d960.E_input', symmetry_r1_xboole_0)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/1fd4e0c15f88414d8c964933c674d960.E_input', t65_xboole_1)).
fof(t65_zfmisc_1, axiom, ![X1, X2]:(k4_xboole_0(X1,k1_tarski(X2))=X1<=>~(r2_hidden(X2,X1))), file('data/MPTP/proofs/1fd4e0c15f88414d8c964933c674d960.E_input', t65_zfmisc_1)).
fof(t90_xboole_1, axiom, ![X1, X2]:r1_xboole_0(k4_xboole_0(X1,k3_xboole_0(X1,X2)),X2), file('data/MPTP/proofs/1fd4e0c15f88414d8c964933c674d960.E_input', t90_xboole_1)).
fof(t70_xboole_1, axiom, ![X1, X2, X3]:(~((~(r1_xboole_0(X1,k2_xboole_0(X2,X3)))&(r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3))))&~((~((r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3)))&r1_xboole_0(X1,k2_xboole_0(X2,X3))))), file('data/MPTP/proofs/1fd4e0c15f88414d8c964933c674d960.E_input', t70_xboole_1)).
fof(c_0_16, plain, ![X18, X19, X20]:((r1_tarski(X18,X19)|~r1_tarski(X18,k4_xboole_0(X19,X20)))&(r1_xboole_0(X18,X20)|~r1_tarski(X18,k4_xboole_0(X19,X20)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t106_xboole_1])])])).
fof(c_0_17, plain, ![X21, X22]:r1_tarski(k4_xboole_0(X21,X22),X21), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_18, plain, ![X27, X28, X30, X31, X32]:(((r2_hidden(esk5_2(X27,X28),X27)|r1_xboole_0(X27,X28))&(r2_hidden(esk5_2(X27,X28),X28)|r1_xboole_0(X27,X28)))&(~r2_hidden(X32,X30)|~r2_hidden(X32,X31)|~r1_xboole_0(X30,X31))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_19, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k4_xboole_0(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_21, plain, ![X41, X42]:(~r1_tarski(X41,X42)|k3_xboole_0(X41,X42)=X41), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_22, plain, ![X9, X10]:k4_xboole_0(X9,k4_xboole_0(X9,X10))=k3_xboole_0(X9,X10), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_23, plain, ![X25, X26]:k3_xboole_0(X25,X26)=k3_xboole_0(X26,X25), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_24, negated_conjecture, ~(![X1, X2, X3, X4]:((r1_tarski(k3_xboole_0(X1,X2),k1_tarski(X4))&(r2_hidden(X4,X3)&r1_xboole_0(X3,X2)))=>r1_xboole_0(k2_xboole_0(X1,X3),X2))), inference(assume_negation,[status(cth)],[t149_zfmisc_1])).
cnf(c_0_25, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, plain, (r1_xboole_0(k4_xboole_0(k4_xboole_0(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
fof(c_0_27, plain, ![X23, X24]:((~r1_xboole_0(X23,X24)|k4_xboole_0(X23,X24)=X23)&(k4_xboole_0(X23,X24)!=X23|r1_xboole_0(X23,X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
cnf(c_0_28, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_31, negated_conjecture, ((r1_tarski(k3_xboole_0(esk1_0,esk2_0),k1_tarski(esk4_0))&(r2_hidden(esk4_0,esk3_0)&r1_xboole_0(esk3_0,esk2_0)))&~r1_xboole_0(k2_xboole_0(esk1_0,esk3_0),esk2_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).
fof(c_0_32, plain, ![X43]:k2_tarski(X43,X43)=k1_tarski(X43), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_33, plain, ![X16, X17]:k1_enumset1(X16,X16,X17)=k2_tarski(X16,X17), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_34, plain, ![X33, X34, X35]:k2_enumset1(X33,X33,X34,X35)=k1_enumset1(X33,X34,X35), inference(variable_rename,[status(thm)],[t71_enumset1])).
cnf(c_0_35, plain, (~r2_hidden(X1,k4_xboole_0(k4_xboole_0(X2,X3),X4))|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_36, plain, (k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=X1|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_38, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_29]), c_0_29])).
fof(c_0_39, plain, ![X37, X38]:(~r1_xboole_0(X37,X38)|r1_xboole_0(X38,X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
fof(c_0_40, plain, ![X36]:r1_xboole_0(X36,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_41, negated_conjecture, (r1_tarski(k3_xboole_0(esk1_0,esk2_0),k1_tarski(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_42, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_44, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45, plain, (~r1_xboole_0(X1,X2)|~r2_hidden(X3,k4_xboole_0(X1,X4))|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_46, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=X2|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_47, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_48, negated_conjecture, (r1_xboole_0(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_49, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_50, plain, (r1_xboole_0(X1,X2)|k4_xboole_0(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_51, negated_conjecture, (r1_tarski(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0)),k2_enumset1(esk4_0,esk4_0,esk4_0,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_42]), c_0_43]), c_0_44]), c_0_29])).
fof(c_0_52, plain, ![X11, X12]:((k4_xboole_0(X11,k1_tarski(X12))!=X11|~r2_hidden(X12,X11))&(r2_hidden(X12,X11)|k4_xboole_0(X11,k1_tarski(X12))=X11)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t65_zfmisc_1])])])).
cnf(c_0_53, plain, (~r1_xboole_0(X1,X2)|~r2_hidden(X3,X4)|~r2_hidden(X3,X2)|~r1_tarski(X4,X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_54, negated_conjecture, (r1_xboole_0(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_55, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=X2|~r1_xboole_0(X2,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_36, c_0_38])).
cnf(c_0_56, plain, (r1_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_47, c_0_49])).
cnf(c_0_57, plain, (r1_xboole_0(X1,k4_xboole_0(X1,X2))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_50, c_0_37])).
cnf(c_0_58, negated_conjecture, (r1_tarski(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk1_0)),k2_enumset1(esk4_0,esk4_0,esk4_0,esk4_0))), inference(rw,[status(thm)],[c_0_51, c_0_38])).
cnf(c_0_59, plain, (r2_hidden(X1,X2)|k4_xboole_0(X2,k1_tarski(X1))=X2), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_60, negated_conjecture, (~r2_hidden(X1,esk3_0)|~r2_hidden(X1,X2)|~r1_tarski(X2,esk2_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_61, negated_conjecture, (r2_hidden(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_62, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k1_xboole_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_63, negated_conjecture, (r1_xboole_0(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk1_0)),k4_xboole_0(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk1_0)),k2_enumset1(esk4_0,esk4_0,esk4_0,esk4_0)))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_64, plain, (k4_xboole_0(X2,k2_enumset1(X1,X1,X1,X1))=X2|r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_42]), c_0_43]), c_0_44])).
cnf(c_0_65, negated_conjecture, (~r2_hidden(esk4_0,X1)|~r1_tarski(X1,esk2_0)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_66, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_36]), c_0_49])])).
cnf(c_0_67, negated_conjecture, (r1_xboole_0(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk1_0)),k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk1_0)))|r2_hidden(esk4_0,k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk1_0)))), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_68, negated_conjecture, (~r2_hidden(esk4_0,k4_xboole_0(esk2_0,X1))), inference(spm,[status(thm)],[c_0_65, c_0_20])).
cnf(c_0_69, plain, (r1_tarski(X1,X1)|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_20, c_0_36])).
fof(c_0_70, plain, ![X39, X40]:r1_xboole_0(k4_xboole_0(X39,k3_xboole_0(X39,X40)),X40), inference(variable_rename,[status(thm)],[t90_xboole_1])).
cnf(c_0_71, plain, (k1_xboole_0=X1|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_36, c_0_66])).
cnf(c_0_72, negated_conjecture, (r1_xboole_0(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk1_0)),k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk1_0)))), inference(sr,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_73, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_69, c_0_49])).
fof(c_0_74, plain, ![X13, X14, X15]:((r1_xboole_0(X13,k2_xboole_0(X14,X15))|(~r1_xboole_0(X13,X14)|~r1_xboole_0(X13,X15)))&((r1_xboole_0(X13,X14)|~r1_xboole_0(X13,k2_xboole_0(X14,X15)))&(r1_xboole_0(X13,X15)|~r1_xboole_0(X13,k2_xboole_0(X14,X15))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t70_xboole_1])])])])).
cnf(c_0_75, plain, (r1_xboole_0(k4_xboole_0(X1,k3_xboole_0(X1,X2)),X2)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_76, negated_conjecture, (k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk1_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_77, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_66]), c_0_73])])).
cnf(c_0_78, plain, (r1_xboole_0(X1,k2_xboole_0(X2,X3))|~r1_xboole_0(X1,X2)|~r1_xboole_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_79, plain, (r1_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X1,X2))),X2)), inference(rw,[status(thm)],[c_0_75, c_0_29])).
cnf(c_0_80, negated_conjecture, (k4_xboole_0(esk2_0,esk1_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_76]), c_0_77]), c_0_20])])).
cnf(c_0_81, negated_conjecture, (~r1_xboole_0(k2_xboole_0(esk1_0,esk3_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_82, plain, (r1_xboole_0(k2_xboole_0(X1,X2),X3)|~r1_xboole_0(X3,X2)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_47, c_0_78])).
cnf(c_0_83, negated_conjecture, (r1_xboole_0(esk2_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_66]), c_0_77])).
cnf(c_0_84, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_54]), c_0_83])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 85
# Proof object clause steps            : 52
# Proof object formula steps           : 33
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 25
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 22
# Processed clauses                    : 853
# ...of these trivial                  : 27
# ...subsumed                          : 559
# ...remaining for further processing  : 267
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 37
# Backward-rewritten                   : 58
# Generated clauses                    : 2655
# ...of the previous two non-trivial   : 1920
# Contextual simplify-reflections      : 5
# Paramodulations                      : 2642
# Factorizations                       : 0
# Equation resolutions                 : 12
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 149
#    Positive orientable unit clauses  : 31
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 108
# Current number of unprocessed clauses: 988
# ...number of literals in the above   : 2393
# Current number of archived formulas  : 0
# Current number of archived clauses   : 122
# Clause-clause subsumption calls (NU) : 2704
# Rec. Clause-clause subsumption calls : 2211
# Non-unit clause-clause subsumptions  : 387
# Unit Clause-clause subsumption calls : 299
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 121
# BW rewrite match successes           : 30
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 26318

# -------------------------------------------------
# User time                : 0.021 s
# System time              : 0.007 s
# Total time               : 0.028 s
# Maximum resident set size: 28636 pages
