# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t54_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&(v2_pre_topc(X3)&l1_pre_topc(X3)))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X1))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X6]:(m1_subset_1(X6,u1_struct_0(X3))=>((r1_tmap_1(X3,X1,X4,X6)&v5_pre_topc(X5,X1,X2))=>r1_tmap_1(X3,X2,k1_partfun1(u1_struct_0(X3),u1_struct_0(X1),u1_struct_0(X1),u1_struct_0(X2),X4,X5),X6)))))))), file('data/MPTP/proofs/202d6ac86107480ca27403ce594a5578.E_input', t54_tmap_1)).
fof(dt_k3_funct_2, axiom, ![X1, X2, X3, X4]:((~(v1_xboole_0(X1))&((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&m1_subset_1(X4,X1)))=>m1_subset_1(k3_funct_2(X1,X2,X3,X4),X2)), file('data/MPTP/proofs/202d6ac86107480ca27403ce594a5578.E_input', dt_k3_funct_2)).
fof(t52_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&(v2_pre_topc(X3)&l1_pre_topc(X3)))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X3))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X1))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))))=>![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>![X7]:(m1_subset_1(X7,u1_struct_0(X3))=>((X7=k3_funct_2(u1_struct_0(X2),u1_struct_0(X3),X4,X6)&(r1_tmap_1(X2,X3,X4,X6)&r1_tmap_1(X3,X1,X5,X7)))=>r1_tmap_1(X2,X1,k1_partfun1(u1_struct_0(X2),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X1),X4,X5),X6))))))))), file('data/MPTP/proofs/202d6ac86107480ca27403ce594a5578.E_input', t52_tmap_1)).
fof(t49_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>(v5_pre_topc(X3,X2,X1)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>r1_tmap_1(X2,X1,X3,X4)))))), file('data/MPTP/proofs/202d6ac86107480ca27403ce594a5578.E_input', t49_tmap_1)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/202d6ac86107480ca27403ce594a5578.E_input', t5_subset)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/202d6ac86107480ca27403ce594a5578.E_input', d1_xboole_0)).
fof(rc7_pre_topc, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>?[X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))&(~(v1_xboole_0(X2))&v4_pre_topc(X2,X1)))), file('data/MPTP/proofs/202d6ac86107480ca27403ce594a5578.E_input', rc7_pre_topc)).
fof(c_0_7, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&(v2_pre_topc(X3)&l1_pre_topc(X3)))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X1))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X6]:(m1_subset_1(X6,u1_struct_0(X3))=>((r1_tmap_1(X3,X1,X4,X6)&v5_pre_topc(X5,X1,X2))=>r1_tmap_1(X3,X2,k1_partfun1(u1_struct_0(X3),u1_struct_0(X1),u1_struct_0(X1),u1_struct_0(X2),X4,X5),X6))))))))), inference(assume_negation,[status(cth)],[t54_tmap_1])).
fof(c_0_8, plain, ![X14, X15, X16, X17]:(v1_xboole_0(X14)|(~v1_funct_1(X16)|(~v1_funct_2(X16,X14,X15)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15))))|~m1_subset_1(X17,X14))|m1_subset_1(k3_funct_2(X14,X15,X16,X17),X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_funct_2])])])).
fof(c_0_9, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&((~v2_struct_0(esk3_0)&(v2_pre_topc(esk3_0)&l1_pre_topc(esk3_0)))&((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,u1_struct_0(esk3_0),u1_struct_0(esk1_0))&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0))))))&((v1_funct_1(esk5_0)&(v1_funct_2(esk5_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))&m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))))&(m1_subset_1(esk6_0,u1_struct_0(esk3_0))&((r1_tmap_1(esk3_0,esk1_0,esk4_0,esk6_0)&v5_pre_topc(esk5_0,esk1_0,esk2_0))&~r1_tmap_1(esk3_0,esk2_0,k1_partfun1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk4_0,esk5_0),esk6_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_7])])])])).
fof(c_0_10, plain, ![X18, X19, X20, X21, X22, X23, X24]:(v2_struct_0(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18))|(v2_struct_0(X19)|(~v2_pre_topc(X19)|~l1_pre_topc(X19))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20))|(~v1_funct_1(X21)|(~v1_funct_2(X21,u1_struct_0(X19),u1_struct_0(X20))|~m1_subset_1(X21,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X19),u1_struct_0(X20)))))|(~v1_funct_1(X22)|(~v1_funct_2(X22,u1_struct_0(X20),u1_struct_0(X18))|~m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X20),u1_struct_0(X18)))))|(~m1_subset_1(X23,u1_struct_0(X19))|(~m1_subset_1(X24,u1_struct_0(X20))|(X24!=k3_funct_2(u1_struct_0(X19),u1_struct_0(X20),X21,X23)|(~r1_tmap_1(X19,X20,X21,X23)|~r1_tmap_1(X20,X18,X22,X24))|r1_tmap_1(X19,X18,k1_partfun1(u1_struct_0(X19),u1_struct_0(X20),u1_struct_0(X20),u1_struct_0(X18),X21,X22),X23))))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t52_tmap_1])])])])).
fof(c_0_11, plain, ![X25, X26, X27, X28]:((~v5_pre_topc(X27,X26,X25)|(~m1_subset_1(X28,u1_struct_0(X26))|r1_tmap_1(X26,X25,X27,X28))|(~v1_funct_1(X27)|(~v1_funct_2(X27,u1_struct_0(X26),u1_struct_0(X25))|~m1_subset_1(X27,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X26),u1_struct_0(X25))))))|(v2_struct_0(X26)|(~v2_pre_topc(X26)|~l1_pre_topc(X26)))|(v2_struct_0(X25)|(~v2_pre_topc(X25)|~l1_pre_topc(X25))))&((m1_subset_1(esk7_3(X25,X26,X27),u1_struct_0(X26))|v5_pre_topc(X27,X26,X25)|(~v1_funct_1(X27)|(~v1_funct_2(X27,u1_struct_0(X26),u1_struct_0(X25))|~m1_subset_1(X27,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X26),u1_struct_0(X25))))))|(v2_struct_0(X26)|(~v2_pre_topc(X26)|~l1_pre_topc(X26)))|(v2_struct_0(X25)|(~v2_pre_topc(X25)|~l1_pre_topc(X25))))&(~r1_tmap_1(X26,X25,X27,esk7_3(X25,X26,X27))|v5_pre_topc(X27,X26,X25)|(~v1_funct_1(X27)|(~v1_funct_2(X27,u1_struct_0(X26),u1_struct_0(X25))|~m1_subset_1(X27,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X26),u1_struct_0(X25))))))|(v2_struct_0(X26)|(~v2_pre_topc(X26)|~l1_pre_topc(X26)))|(v2_struct_0(X25)|(~v2_pre_topc(X25)|~l1_pre_topc(X25)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t49_tmap_1])])])])])])).
cnf(c_0_12, plain, (v1_xboole_0(X1)|m1_subset_1(k3_funct_2(X1,X3,X2,X4),X3)|~v1_funct_1(X2)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, negated_conjecture, (v1_funct_2(esk4_0,u1_struct_0(esk3_0),u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_16, plain, ![X30, X31, X32]:(~r2_hidden(X30,X31)|(~m1_subset_1(X31,k1_zfmisc_1(X32))|~v1_xboole_0(X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_17, plain, ![X33, X34, X35]:((~v1_xboole_0(X33)|~r2_hidden(X34,X33))&(r2_hidden(esk8_1(X35),X35)|v1_xboole_0(X35))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_18, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|r1_tmap_1(X2,X1,k1_partfun1(u1_struct_0(X2),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X1),X4,X5),X6)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~v1_funct_1(X4)|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X1))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X3))|X7!=k3_funct_2(u1_struct_0(X2),u1_struct_0(X3),X4,X6)|~r1_tmap_1(X2,X3,X4,X6)|~r1_tmap_1(X3,X1,X5,X7)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19, plain, (r1_tmap_1(X2,X3,X1,X4)|v2_struct_0(X2)|v2_struct_0(X3)|~v5_pre_topc(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20, negated_conjecture, (v5_pre_topc(esk5_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_21, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_22, negated_conjecture, (v1_funct_2(esk5_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_23, negated_conjecture, (v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_24, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_25, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_26, negated_conjecture, (v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_27, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_28, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_29, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_30, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|m1_subset_1(k3_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk1_0),esk4_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_15])])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk6_0,u1_struct_0(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_32, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_33, plain, (r2_hidden(esk8_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_34, plain, ![X37]:((m1_subset_1(esk9_1(X37),k1_zfmisc_1(u1_struct_0(X37)))|(v2_struct_0(X37)|(~v2_pre_topc(X37)|~l1_pre_topc(X37))))&((~v1_xboole_0(esk9_1(X37))|(v2_struct_0(X37)|(~v2_pre_topc(X37)|~l1_pre_topc(X37))))&(v4_pre_topc(esk9_1(X37),X37)|(v2_struct_0(X37)|(~v2_pre_topc(X37)|~l1_pre_topc(X37)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[rc7_pre_topc])])])])])).
cnf(c_0_35, plain, (r1_tmap_1(X1,X2,k1_partfun1(u1_struct_0(X1),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X2),X4,X5),X6)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~r1_tmap_1(X3,X2,X5,k3_funct_2(u1_struct_0(X1),u1_struct_0(X3),X4,X6))|~r1_tmap_1(X1,X3,X4,X6)|~m1_subset_1(k3_funct_2(u1_struct_0(X1),u1_struct_0(X3),X4,X6),u1_struct_0(X3))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~m1_subset_1(X6,u1_struct_0(X1))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(er,[status(thm)],[c_0_18])).
cnf(c_0_36, negated_conjecture, (r1_tmap_1(esk3_0,esk1_0,esk4_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_37, negated_conjecture, (l1_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_38, negated_conjecture, (v2_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_39, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_40, negated_conjecture, (r1_tmap_1(esk1_0,esk2_0,esk5_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27])]), c_0_28]), c_0_29])).
cnf(c_0_41, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|m1_subset_1(k3_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk1_0),esk4_0,esk6_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_42, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_43, plain, (m1_subset_1(esk9_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_44, plain, (v2_struct_0(X1)|~v1_xboole_0(esk9_1(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45, negated_conjecture, (r1_tmap_1(esk3_0,X1,k1_partfun1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(X1),esk4_0,X2),esk6_0)|v2_struct_0(X1)|~r1_tmap_1(esk1_0,X1,X2,k3_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk1_0),esk4_0,esk6_0))|~m1_subset_1(k3_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk1_0),esk4_0,esk6_0),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(esk1_0),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_14]), c_0_31]), c_0_13]), c_0_15]), c_0_25]), c_0_37]), c_0_27]), c_0_38])]), c_0_39]), c_0_29])).
cnf(c_0_46, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_tmap_1(esk1_0,esk2_0,esk5_0,k3_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk1_0),esk4_0,esk6_0))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_47, negated_conjecture, (~r1_tmap_1(esk3_0,esk2_0,k1_partfun1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk4_0,esk5_0),esk6_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_48, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])).
cnf(c_0_49, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_26])]), c_0_47]), c_0_28]), c_0_41])).
cnf(c_0_50, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_37]), c_0_38])]), c_0_39]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 51
# Proof object clause steps            : 36
# Proof object formula steps           : 15
# Proof object conjectures             : 29
# Proof object clause conjectures      : 26
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 39
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 75
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 75
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 5
# Generated clauses                    : 22
# ...of the previous two non-trivial   : 19
# Contextual simplify-reflections      : 2
# Paramodulations                      : 21
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 39
#    Positive orientable unit clauses  : 16
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 19
# Current number of unprocessed clauses: 4
# ...number of literals in the above   : 30
# Current number of archived formulas  : 0
# Current number of archived clauses   : 35
# Clause-clause subsumption calls (NU) : 1265
# Rec. Clause-clause subsumption calls : 59
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 6
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4028

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.003 s
# Total time               : 0.011 s
# Maximum resident set size: 28792 pages
