# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t17_subset_1, conjecture, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(X1))=>(![X5]:(m1_subset_1(X5,X1)=>(r2_hidden(X5,X2)<=>(r2_hidden(X5,X3)&~(r2_hidden(X5,X4)))))=>X2=k7_subset_1(X1,X3,X4))))), file('data/MPTP/proofs/20b11bb10ae74192b8c4d11a93d75fff.E_input', t17_subset_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/20b11bb10ae74192b8c4d11a93d75fff.E_input', d2_subset_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/20b11bb10ae74192b8c4d11a93d75fff.E_input', d1_xboole_0)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/20b11bb10ae74192b8c4d11a93d75fff.E_input', d1_zfmisc_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/20b11bb10ae74192b8c4d11a93d75fff.E_input', d3_tarski)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/20b11bb10ae74192b8c4d11a93d75fff.E_input', fc1_subset_1)).
fof(redefinition_k7_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k7_subset_1(X1,X2,X3)=k4_xboole_0(X2,X3)), file('data/MPTP/proofs/20b11bb10ae74192b8c4d11a93d75fff.E_input', redefinition_k7_subset_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/20b11bb10ae74192b8c4d11a93d75fff.E_input', t100_xboole_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/20b11bb10ae74192b8c4d11a93d75fff.E_input', d5_xboole_0)).
fof(c_0_9, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(X1))=>(![X5]:(m1_subset_1(X5,X1)=>(r2_hidden(X5,X2)<=>(r2_hidden(X5,X3)&~(r2_hidden(X5,X4)))))=>X2=k7_subset_1(X1,X3,X4)))))), inference(assume_negation,[status(cth)],[t17_subset_1])).
fof(c_0_10, plain, ![X13, X14]:(((~m1_subset_1(X14,X13)|r2_hidden(X14,X13)|v1_xboole_0(X13))&(~r2_hidden(X14,X13)|m1_subset_1(X14,X13)|v1_xboole_0(X13)))&((~m1_subset_1(X14,X13)|v1_xboole_0(X14)|~v1_xboole_0(X13))&(~v1_xboole_0(X14)|m1_subset_1(X14,X13)|~v1_xboole_0(X13)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_11, plain, ![X27, X28, X29]:((~v1_xboole_0(X27)|~r2_hidden(X28,X27))&(r2_hidden(esk6_1(X29),X29)|v1_xboole_0(X29))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_12, plain, ![X38, X39, X40, X41, X42, X43]:(((~r2_hidden(X40,X39)|r1_tarski(X40,X38)|X39!=k1_zfmisc_1(X38))&(~r1_tarski(X41,X38)|r2_hidden(X41,X39)|X39!=k1_zfmisc_1(X38)))&((~r2_hidden(esk8_2(X42,X43),X43)|~r1_tarski(esk8_2(X42,X43),X42)|X43=k1_zfmisc_1(X42))&(r2_hidden(esk8_2(X42,X43),X43)|r1_tarski(esk8_2(X42,X43),X42)|X43=k1_zfmisc_1(X42)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
fof(c_0_13, negated_conjecture, ![X10]:(m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk4_0,k1_zfmisc_1(esk1_0))&((((r2_hidden(X10,esk3_0)|~r2_hidden(X10,esk2_0)|~m1_subset_1(X10,esk1_0))&(~r2_hidden(X10,esk4_0)|~r2_hidden(X10,esk2_0)|~m1_subset_1(X10,esk1_0)))&(~r2_hidden(X10,esk3_0)|r2_hidden(X10,esk4_0)|r2_hidden(X10,esk2_0)|~m1_subset_1(X10,esk1_0)))&esk2_0!=k7_subset_1(esk1_0,esk3_0,esk4_0))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])])])).
cnf(c_0_14, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_16, plain, ![X31, X32, X33, X34, X35]:((~r1_tarski(X31,X32)|(~r2_hidden(X33,X31)|r2_hidden(X33,X32)))&((r2_hidden(esk7_2(X34,X35),X34)|r1_tarski(X34,X35))&(~r2_hidden(esk7_2(X34,X35),X35)|r1_tarski(X34,X35)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_17, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_18, plain, ![X37]:~v1_xboole_0(k1_zfmisc_1(X37)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
fof(c_0_19, plain, ![X15, X16, X17]:(~m1_subset_1(X16,k1_zfmisc_1(X15))|k7_subset_1(X15,X16,X17)=k4_xboole_0(X16,X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_subset_1])])).
fof(c_0_20, plain, ![X11, X12]:k4_xboole_0(X11,X12)=k5_xboole_0(X11,k3_xboole_0(X11,X12)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_21, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_23, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_26, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_27, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, plain, (k7_subset_1(X2,X1,X3)=k4_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30, plain, (r1_tarski(X1,X2)|~r2_hidden(esk7_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_31, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_32, plain, (r2_hidden(X1,X2)|~r2_hidden(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_33, negated_conjecture, (r2_hidden(esk2_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])).
cnf(c_0_34, plain, (k7_subset_1(X2,X1,X3)=k5_xboole_0(X1,k3_xboole_0(X1,X3))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35, negated_conjecture, (r1_tarski(X1,esk3_0)|~r2_hidden(esk7_2(X1,esk3_0),esk2_0)|~r2_hidden(esk7_2(X1,esk3_0),esk1_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_36, plain, (r2_hidden(esk7_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_37, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_38, negated_conjecture, (esk2_0!=k7_subset_1(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_39, plain, (k7_subset_1(X1,X2,X3)=k7_subset_1(X4,X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X4))), inference(spm,[status(thm)],[c_0_34, c_0_34])).
cnf(c_0_40, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_41, plain, ![X18, X19, X20, X21, X22, X23, X24, X25]:((((r2_hidden(X21,X18)|~r2_hidden(X21,X20)|X20!=k4_xboole_0(X18,X19))&(~r2_hidden(X21,X19)|~r2_hidden(X21,X20)|X20!=k4_xboole_0(X18,X19)))&(~r2_hidden(X22,X18)|r2_hidden(X22,X19)|r2_hidden(X22,X20)|X20!=k4_xboole_0(X18,X19)))&((~r2_hidden(esk5_3(X23,X24,X25),X25)|(~r2_hidden(esk5_3(X23,X24,X25),X23)|r2_hidden(esk5_3(X23,X24,X25),X24))|X25=k4_xboole_0(X23,X24))&((r2_hidden(esk5_3(X23,X24,X25),X23)|r2_hidden(esk5_3(X23,X24,X25),X25)|X25=k4_xboole_0(X23,X24))&(~r2_hidden(esk5_3(X23,X24,X25),X24)|r2_hidden(esk5_3(X23,X24,X25),X25)|X25=k4_xboole_0(X23,X24))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_42, negated_conjecture, (r1_tarski(esk2_0,esk3_0)|~r2_hidden(esk7_2(esk2_0,esk3_0),esk1_0)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_43, negated_conjecture, (r1_tarski(esk2_0,X1)|r2_hidden(esk7_2(esk2_0,X1),esk1_0)), inference(spm,[status(thm)],[c_0_37, c_0_36])).
cnf(c_0_44, negated_conjecture, (k7_subset_1(X1,esk3_0,esk4_0)!=esk2_0|~m1_subset_1(esk3_0,k1_zfmisc_1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])])).
cnf(c_0_45, plain, (r2_hidden(esk5_3(X1,X2,X3),X1)|r2_hidden(esk5_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_46, negated_conjecture, (r1_tarski(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_47, plain, (r2_hidden(esk5_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk5_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_48, negated_conjecture, (r2_hidden(X1,esk4_0)|r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_49, negated_conjecture, (r2_hidden(esk3_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_40]), c_0_27])).
cnf(c_0_50, negated_conjecture, (k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,esk4_0))!=esk2_0|~m1_subset_1(esk3_0,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_44, c_0_34])).
cnf(c_0_51, plain, (X3=k5_xboole_0(X1,k3_xboole_0(X1,X2))|r2_hidden(esk5_3(X1,X2,X3),X3)|r2_hidden(esk5_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_45, c_0_29])).
cnf(c_0_52, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_23, c_0_46])).
cnf(c_0_53, plain, (X3=k5_xboole_0(X1,k3_xboole_0(X1,X2))|r2_hidden(esk5_3(X1,X2,X3),X3)|~r2_hidden(esk5_3(X1,X2,X3),X2)), inference(rw,[status(thm)],[c_0_47, c_0_29])).
cnf(c_0_54, negated_conjecture, (r2_hidden(X1,esk4_0)|r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_48, c_0_22])).
cnf(c_0_55, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_32, c_0_49])).
cnf(c_0_56, negated_conjecture, (r2_hidden(esk5_3(esk3_0,esk4_0,esk2_0),esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(X1))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51])]), c_0_52])).
cnf(c_0_57, negated_conjecture, (X1=k5_xboole_0(X2,k3_xboole_0(X2,esk4_0))|r2_hidden(esk5_3(X2,esk4_0,X1),esk2_0)|r2_hidden(esk5_3(X2,esk4_0,X1),X1)|~r2_hidden(esk5_3(X2,esk4_0,X1),esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55])).
cnf(c_0_58, negated_conjecture, (r2_hidden(esk5_3(esk3_0,esk4_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_56, c_0_40])).
cnf(c_0_59, negated_conjecture, (k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,esk4_0))=esk2_0|r2_hidden(esk5_3(esk3_0,esk4_0,esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_60, plain, (r2_hidden(esk5_3(X1,X2,X3),X2)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk5_3(X1,X2,X3),X3)|~r2_hidden(esk5_3(X1,X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_61, negated_conjecture, (r2_hidden(esk5_3(esk3_0,esk4_0,esk2_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_50, c_0_59])).
cnf(c_0_62, plain, (X3=k5_xboole_0(X1,k3_xboole_0(X1,X2))|r2_hidden(esk5_3(X1,X2,X3),X2)|~r2_hidden(esk5_3(X1,X2,X3),X3)|~r2_hidden(esk5_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_60, c_0_29])).
cnf(c_0_63, negated_conjecture, (r2_hidden(esk5_3(esk3_0,esk4_0,esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_61, c_0_40])).
cnf(c_0_64, negated_conjecture, (k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,esk4_0))=esk2_0|r2_hidden(esk5_3(esk3_0,esk4_0,esk2_0),esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_58])])).
cnf(c_0_65, negated_conjecture, (~r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_66, negated_conjecture, (r2_hidden(esk5_3(esk3_0,esk4_0,esk2_0),esk4_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_50, c_0_64])).
cnf(c_0_67, negated_conjecture, (~r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_65, c_0_22])).
cnf(c_0_68, negated_conjecture, (r2_hidden(esk5_3(esk3_0,esk4_0,esk2_0),esk4_0)), inference(spm,[status(thm)],[c_0_66, c_0_40])).
cnf(c_0_69, negated_conjecture, (r2_hidden(esk5_3(esk3_0,esk4_0,esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_55, c_0_58])).
cnf(c_0_70, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_63]), c_0_69])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 71
# Proof object clause steps            : 52
# Proof object formula steps           : 19
# Proof object conjectures             : 34
# Proof object clause conjectures      : 31
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 28
# Processed clauses                    : 554
# ...of these trivial                  : 27
# ...subsumed                          : 213
# ...remaining for further processing  : 314
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 21
# Backward-rewritten                   : 20
# Generated clauses                    : 3335
# ...of the previous two non-trivial   : 3044
# Contextual simplify-reflections      : 12
# Paramodulations                      : 3300
# Factorizations                       : 18
# Equation resolutions                 : 12
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 235
#    Positive orientable unit clauses  : 51
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 179
# Current number of unprocessed clauses: 2493
# ...number of literals in the above   : 10505
# Current number of archived formulas  : 0
# Current number of archived clauses   : 75
# Clause-clause subsumption calls (NU) : 6133
# Rec. Clause-clause subsumption calls : 3990
# Non-unit clause-clause subsumptions  : 154
# Unit Clause-clause subsumption calls : 1227
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 111
# BW rewrite match successes           : 15
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 55500

# -------------------------------------------------
# User time                : 0.046 s
# System time              : 0.006 s
# Total time               : 0.052 s
# Maximum resident set size: 28680 pages
