# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_funct_2, axiom, ![X1, X2, X3]:(X3=k1_funct_2(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:((v1_relat_1(X5)&v1_funct_1(X5))&(X4=X5&(k1_relat_1(X5)=X1&r1_tarski(k2_relat_1(X5),X2)))))), file('data/MPTP/proofs/20e0de8865c74abdbef50cd8c5224bc8.E_input', d2_funct_2)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/20e0de8865c74abdbef50cd8c5224bc8.E_input', t5_subset)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/20e0de8865c74abdbef50cd8c5224bc8.E_input', t3_subset)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/20e0de8865c74abdbef50cd8c5224bc8.E_input', d10_xboole_0)).
fof(t121_funct_2, conjecture, ![X1, X2, X3]:(r2_hidden(X3,k1_funct_2(X1,X2))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))), file('data/MPTP/proofs/20e0de8865c74abdbef50cd8c5224bc8.E_input', t121_funct_2)).
fof(d5_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4))))), file('data/MPTP/proofs/20e0de8865c74abdbef50cd8c5224bc8.E_input', d5_funct_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('data/MPTP/proofs/20e0de8865c74abdbef50cd8c5224bc8.E_input', cc2_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/20e0de8865c74abdbef50cd8c5224bc8.E_input', cc1_relset_1)).
fof(t3_funct_2, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_funct_1(X1)&(v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))))), file('data/MPTP/proofs/20e0de8865c74abdbef50cd8c5224bc8.E_input', t3_funct_2)).
fof(cc4_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/20e0de8865c74abdbef50cd8c5224bc8.E_input', cc4_relset_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/20e0de8865c74abdbef50cd8c5224bc8.E_input', d3_tarski)).
fof(dt_k6_partfun1, axiom, ![X1]:(v1_partfun1(k6_partfun1(X1),X1)&m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), file('data/MPTP/proofs/20e0de8865c74abdbef50cd8c5224bc8.E_input', dt_k6_partfun1)).
fof(cc5_funct_2, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_funct_2(X3,X1,X2))=>(v1_funct_1(X3)&v1_partfun1(X3,X1))))), file('data/MPTP/proofs/20e0de8865c74abdbef50cd8c5224bc8.E_input', cc5_funct_2)).
fof(d19_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(v5_relat_1(X2,X1)<=>r1_tarski(k2_relat_1(X2),X1))), file('data/MPTP/proofs/20e0de8865c74abdbef50cd8c5224bc8.E_input', d19_relat_1)).
fof(t11_relset_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>((r1_tarski(k1_relat_1(X3),X1)&r1_tarski(k2_relat_1(X3),X2))=>m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))), file('data/MPTP/proofs/20e0de8865c74abdbef50cd8c5224bc8.E_input', t11_relset_1)).
fof(cc1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_partfun1(X3,X1))=>(v1_funct_1(X3)&v1_funct_2(X3,X1,X2)))), file('data/MPTP/proofs/20e0de8865c74abdbef50cd8c5224bc8.E_input', cc1_funct_2)).
fof(c_0_16, plain, ![X45, X46, X47, X48, X50, X51, X52, X53, X54, X56]:(((((v1_relat_1(esk8_4(X45,X46,X47,X48))|~r2_hidden(X48,X47)|X47!=k1_funct_2(X45,X46))&(v1_funct_1(esk8_4(X45,X46,X47,X48))|~r2_hidden(X48,X47)|X47!=k1_funct_2(X45,X46)))&((X48=esk8_4(X45,X46,X47,X48)|~r2_hidden(X48,X47)|X47!=k1_funct_2(X45,X46))&((k1_relat_1(esk8_4(X45,X46,X47,X48))=X45|~r2_hidden(X48,X47)|X47!=k1_funct_2(X45,X46))&(r1_tarski(k2_relat_1(esk8_4(X45,X46,X47,X48)),X46)|~r2_hidden(X48,X47)|X47!=k1_funct_2(X45,X46)))))&(~v1_relat_1(X51)|~v1_funct_1(X51)|(X50!=X51|(k1_relat_1(X51)!=X45|~r1_tarski(k2_relat_1(X51),X46)))|r2_hidden(X50,X47)|X47!=k1_funct_2(X45,X46)))&((~r2_hidden(esk9_3(X52,X53,X54),X54)|(~v1_relat_1(X56)|~v1_funct_1(X56)|(esk9_3(X52,X53,X54)!=X56|(k1_relat_1(X56)!=X52|~r1_tarski(k2_relat_1(X56),X53))))|X54=k1_funct_2(X52,X53))&(((v1_relat_1(esk10_3(X52,X53,X54))|r2_hidden(esk9_3(X52,X53,X54),X54)|X54=k1_funct_2(X52,X53))&(v1_funct_1(esk10_3(X52,X53,X54))|r2_hidden(esk9_3(X52,X53,X54),X54)|X54=k1_funct_2(X52,X53)))&((esk9_3(X52,X53,X54)=esk10_3(X52,X53,X54)|r2_hidden(esk9_3(X52,X53,X54),X54)|X54=k1_funct_2(X52,X53))&((k1_relat_1(esk10_3(X52,X53,X54))=X52|r2_hidden(esk9_3(X52,X53,X54),X54)|X54=k1_funct_2(X52,X53))&(r1_tarski(k2_relat_1(esk10_3(X52,X53,X54)),X53)|r2_hidden(esk9_3(X52,X53,X54),X54)|X54=k1_funct_2(X52,X53))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_funct_2])])])])])])).
fof(c_0_17, plain, ![X32, X33, X34]:(~r2_hidden(X32,X33)|(~m1_subset_1(X33,k1_zfmisc_1(X34))|~v1_xboole_0(X34))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_18, plain, ![X12, X13]:((~m1_subset_1(X12,k1_zfmisc_1(X13))|r1_tarski(X12,X13))&(~r1_tarski(X12,X13)|m1_subset_1(X12,k1_zfmisc_1(X13)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_19, plain, ![X59, X60]:(((r1_tarski(X59,X60)|X59!=X60)&(r1_tarski(X60,X59)|X59!=X60))&(~r1_tarski(X59,X60)|~r1_tarski(X60,X59)|X59=X60)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_20, plain, (k1_relat_1(esk8_4(X1,X2,X3,X4))=X1|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, plain, (X1=esk8_4(X2,X3,X4,X1)|~r2_hidden(X1,X4)|X4!=k1_funct_2(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, negated_conjecture, ~(![X1, X2, X3]:(r2_hidden(X3,k1_funct_2(X1,X2))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))), inference(assume_negation,[status(cth)],[t121_funct_2])).
cnf(c_0_23, plain, (v1_relat_1(esk8_4(X1,X2,X3,X4))|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (v1_funct_1(esk8_4(X1,X2,X3,X4))|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_28, plain, ![X35, X36, X37, X39, X40, X41, X43]:((((r2_hidden(esk5_3(X35,X36,X37),k1_relat_1(X35))|~r2_hidden(X37,X36)|X36!=k2_relat_1(X35)|(~v1_relat_1(X35)|~v1_funct_1(X35)))&(X37=k1_funct_1(X35,esk5_3(X35,X36,X37))|~r2_hidden(X37,X36)|X36!=k2_relat_1(X35)|(~v1_relat_1(X35)|~v1_funct_1(X35))))&(~r2_hidden(X40,k1_relat_1(X35))|X39!=k1_funct_1(X35,X40)|r2_hidden(X39,X36)|X36!=k2_relat_1(X35)|(~v1_relat_1(X35)|~v1_funct_1(X35))))&((~r2_hidden(esk6_2(X35,X41),X41)|(~r2_hidden(X43,k1_relat_1(X35))|esk6_2(X35,X41)!=k1_funct_1(X35,X43))|X41=k2_relat_1(X35)|(~v1_relat_1(X35)|~v1_funct_1(X35)))&((r2_hidden(esk7_2(X35,X41),k1_relat_1(X35))|r2_hidden(esk6_2(X35,X41),X41)|X41=k2_relat_1(X35)|(~v1_relat_1(X35)|~v1_funct_1(X35)))&(esk6_2(X35,X41)=k1_funct_1(X35,esk7_2(X35,X41))|r2_hidden(esk6_2(X35,X41),X41)|X41=k2_relat_1(X35)|(~v1_relat_1(X35)|~v1_funct_1(X35)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_funct_1])])])])])])).
cnf(c_0_29, plain, (k1_relat_1(esk8_4(X1,X2,k1_funct_2(X1,X2),X3))=X1|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_30, plain, (esk8_4(X1,X2,k1_funct_2(X1,X2),X3)=X3|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_21])).
fof(c_0_31, negated_conjecture, (r2_hidden(esk3_0,k1_funct_2(esk1_0,esk2_0))&(~v1_funct_1(esk3_0)|(~v1_funct_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])).
cnf(c_0_32, plain, (v1_relat_1(esk8_4(X1,X2,k1_funct_2(X1,X2),X3))|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_23])).
cnf(c_0_33, plain, (v1_funct_1(esk8_4(X1,X2,k1_funct_2(X1,X2),X3))|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_24])).
fof(c_0_34, plain, ![X61, X62, X63]:((v4_relat_1(X63,X61)|~m1_subset_1(X63,k1_zfmisc_1(k2_zfmisc_1(X61,X62))))&(v5_relat_1(X63,X62)|~m1_subset_1(X63,k1_zfmisc_1(k2_zfmisc_1(X61,X62))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_35, plain, ![X17, X18, X19]:(~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(X17,X18)))|v1_relat_1(X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_36, plain, (~r1_tarski(X1,X2)|~v1_xboole_0(X2)|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_37, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_38, plain, (r2_hidden(X3,X4)|~r2_hidden(X1,k1_relat_1(X2))|X3!=k1_funct_1(X2,X1)|X4!=k2_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_39, plain, ![X20]:((v1_funct_1(X20)|(~v1_relat_1(X20)|~v1_funct_1(X20)))&((v1_funct_2(X20,k1_relat_1(X20),k2_relat_1(X20))|(~v1_relat_1(X20)|~v1_funct_1(X20)))&(m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X20),k2_relat_1(X20))))|(~v1_relat_1(X20)|~v1_funct_1(X20))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_funct_2])])])).
cnf(c_0_40, plain, (k1_relat_1(X1)=X2|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_41, negated_conjecture, (r2_hidden(esk3_0,k1_funct_2(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_42, plain, (v1_relat_1(X1)|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(spm,[status(thm)],[c_0_32, c_0_30])).
cnf(c_0_43, plain, (v1_funct_1(X1)|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(spm,[status(thm)],[c_0_33, c_0_30])).
fof(c_0_44, plain, ![X9, X10, X11]:(~v1_xboole_0(X9)|(~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X10,X9)))|v1_xboole_0(X11))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])).
cnf(c_0_45, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_46, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_47, plain, ![X23, X24, X25, X26, X27]:((~r1_tarski(X23,X24)|(~r2_hidden(X25,X23)|r2_hidden(X25,X24)))&((r2_hidden(esk4_2(X26,X27),X26)|r1_tarski(X26,X27))&(~r2_hidden(esk4_2(X26,X27),X27)|r1_tarski(X26,X27)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_48, plain, ![X58]:(v1_partfun1(k6_partfun1(X58),X58)&m1_subset_1(k6_partfun1(X58),k1_zfmisc_1(k2_zfmisc_1(X58,X58)))), inference(variable_rename,[status(thm)],[dt_k6_partfun1])).
cnf(c_0_49, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_50, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_38])])).
fof(c_0_51, plain, ![X64, X65, X66]:((v1_funct_1(X66)|(~v1_funct_1(X66)|~v1_funct_2(X66,X64,X65))|~m1_subset_1(X66,k1_zfmisc_1(k2_zfmisc_1(X64,X65)))|v1_xboole_0(X65))&(v1_partfun1(X66,X64)|(~v1_funct_1(X66)|~v1_funct_2(X66,X64,X65))|~m1_subset_1(X66,k1_zfmisc_1(k2_zfmisc_1(X64,X65)))|v1_xboole_0(X65))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_funct_2])])])])])).
cnf(c_0_52, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_53, negated_conjecture, (k1_relat_1(esk3_0)=esk1_0), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_54, negated_conjecture, (v1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_42, c_0_41])).
cnf(c_0_55, negated_conjecture, (v1_funct_1(esk3_0)), inference(spm,[status(thm)],[c_0_43, c_0_41])).
cnf(c_0_56, plain, (v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_57, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_58, plain, ![X21, X22]:((~v5_relat_1(X22,X21)|r1_tarski(k2_relat_1(X22),X21)|~v1_relat_1(X22))&(~r1_tarski(k2_relat_1(X22),X21)|v5_relat_1(X22,X21)|~v1_relat_1(X22))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d19_relat_1])])])).
cnf(c_0_59, plain, (v5_relat_1(X1,X2)|~r1_tarski(X1,k2_zfmisc_1(X3,X2))), inference(spm,[status(thm)],[c_0_45, c_0_26])).
cnf(c_0_60, plain, (v1_relat_1(X1)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_46, c_0_26])).
cnf(c_0_61, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_62, plain, (m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_63, plain, (~v1_relat_1(X1)|~v1_xboole_0(k2_relat_1(X1))|~v1_funct_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_64, plain, (v1_partfun1(X1,X2)|v1_xboole_0(X3)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_65, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54]), c_0_55])])).
cnf(c_0_66, negated_conjecture, (v1_funct_2(esk3_0,esk1_0,k2_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_53]), c_0_54]), c_0_55])])).
fof(c_0_67, plain, ![X14, X15, X16]:(~v1_relat_1(X16)|(~r1_tarski(k1_relat_1(X16),X14)|~r1_tarski(k2_relat_1(X16),X15)|m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_relset_1])])).
cnf(c_0_68, plain, (~v1_relat_1(X1)|~v1_xboole_0(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))|~v1_funct_1(X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_25, c_0_52])).
cnf(c_0_69, plain, (v1_xboole_0(X1)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))|~v1_xboole_0(X3)), inference(spm,[status(thm)],[c_0_57, c_0_26])).
cnf(c_0_70, plain, (r1_tarski(k2_relat_1(X1),X2)|~v5_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_71, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2)), inference(spm,[status(thm)],[c_0_59, c_0_37])).
cnf(c_0_72, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_60, c_0_37])).
cnf(c_0_73, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_49, c_0_61])).
cnf(c_0_74, plain, (v1_xboole_0(k6_partfun1(X1))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_57, c_0_62])).
cnf(c_0_75, negated_conjecture, (~v1_xboole_0(k2_relat_1(esk3_0))|~r2_hidden(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_53]), c_0_54]), c_0_55])])).
cnf(c_0_76, negated_conjecture, (v1_partfun1(esk3_0,esk1_0)|v1_xboole_0(k2_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_66]), c_0_55])])).
cnf(c_0_77, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_78, negated_conjecture, (~v1_xboole_0(k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0)))|~r2_hidden(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_53]), c_0_54]), c_0_55])])).
cnf(c_0_79, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_69, c_0_37])).
cnf(c_0_80, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_81, plain, (r1_tarski(k2_relat_1(k2_zfmisc_1(X1,X2)),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_72])])).
cnf(c_0_82, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_83, plain, (r1_tarski(k6_partfun1(X1),X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_84, negated_conjecture, (v1_partfun1(esk3_0,esk1_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_85, plain, (~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X3)|~v1_xboole_0(k2_zfmisc_1(X2,X3))|~r2_hidden(X4,X1)), inference(spm,[status(thm)],[c_0_25, c_0_77])).
cnf(c_0_86, negated_conjecture, (~v1_xboole_0(k2_relat_1(esk3_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_87, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k2_relat_1(k2_zfmisc_1(X3,X2)))), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_88, plain, (X1=k6_partfun1(X2)|~r1_tarski(X1,k6_partfun1(X2))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_89, negated_conjecture, (v1_partfun1(esk3_0,esk1_0)|r1_tarski(esk1_0,X1)), inference(spm,[status(thm)],[c_0_84, c_0_61])).
cnf(c_0_90, plain, (~v1_relat_1(X1)|~r1_tarski(k2_relat_1(X1),X2)|~v1_xboole_0(k2_zfmisc_1(k1_relat_1(X1),X2))|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_85, c_0_37])).
cnf(c_0_91, negated_conjecture, (v1_partfun1(esk3_0,esk1_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_86, c_0_76])).
cnf(c_0_92, plain, (r1_tarski(k2_relat_1(k2_zfmisc_1(X1,X2)),X3)|r2_hidden(esk4_2(k2_relat_1(k2_zfmisc_1(X1,X2)),X3),X2)), inference(spm,[status(thm)],[c_0_87, c_0_61])).
cnf(c_0_93, negated_conjecture, (k6_partfun1(X1)=esk1_0|v1_partfun1(esk3_0,esk1_0)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_94, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_95, plain, (~v1_relat_1(X1)|~r1_tarski(k2_relat_1(X1),X2)|~v1_xboole_0(X2)|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_90, c_0_79])).
cnf(c_0_96, negated_conjecture, (X1=esk1_0|v1_partfun1(esk3_0,esk1_0)|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_82, c_0_89])).
cnf(c_0_97, negated_conjecture, (v1_partfun1(esk3_0,esk1_0)|r1_tarski(k2_relat_1(k2_zfmisc_1(X1,esk3_0)),X2)), inference(spm,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_98, negated_conjecture, (v1_partfun1(esk3_0,esk1_0)|v1_xboole_0(esk1_0)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_74, c_0_93])).
cnf(c_0_99, plain, (r1_tarski(k6_partfun1(X1),k2_zfmisc_1(X1,X1))), inference(spm,[status(thm)],[c_0_94, c_0_62])).
cnf(c_0_100, plain, (~v1_relat_1(X1)|~v1_xboole_0(k2_relat_1(X1))|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_95, c_0_37])).
cnf(c_0_101, negated_conjecture, (k2_relat_1(k2_zfmisc_1(X1,esk3_0))=esk1_0|v1_partfun1(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_102, negated_conjecture, (v1_partfun1(esk3_0,esk1_0)|v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_98, c_0_76])).
cnf(c_0_103, plain, (r2_hidden(X1,k2_zfmisc_1(X2,X2))|~r2_hidden(X1,k6_partfun1(X2))), inference(spm,[status(thm)],[c_0_80, c_0_99])).
cnf(c_0_104, negated_conjecture, (v1_partfun1(esk3_0,esk1_0)|~r2_hidden(X1,k2_zfmisc_1(X2,esk3_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_72])]), c_0_102])).
cnf(c_0_105, plain, (r1_tarski(k6_partfun1(X1),X2)|r2_hidden(esk4_2(k6_partfun1(X1),X2),k2_zfmisc_1(X1,X1))), inference(spm,[status(thm)],[c_0_103, c_0_61])).
cnf(c_0_106, negated_conjecture, (v1_partfun1(esk3_0,esk1_0)|r1_tarski(esk3_0,X1)), inference(spm,[status(thm)],[c_0_91, c_0_61])).
cnf(c_0_107, negated_conjecture, (v1_partfun1(esk3_0,esk1_0)|r1_tarski(k6_partfun1(esk3_0),X1)), inference(spm,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_108, plain, (v1_partfun1(k6_partfun1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_109, negated_conjecture, (k6_partfun1(X1)=esk3_0|v1_partfun1(esk3_0,esk1_0)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_88, c_0_106])).
cnf(c_0_110, negated_conjecture, (v1_partfun1(esk3_0,esk1_0)|v1_xboole_0(k6_partfun1(esk3_0))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_69, c_0_107])).
fof(c_0_111, plain, ![X29, X30, X31]:((v1_funct_1(X31)|(~v1_funct_1(X31)|~v1_partfun1(X31,X29))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(X29,X30))))&(v1_funct_2(X31,X29,X30)|(~v1_funct_1(X31)|~v1_partfun1(X31,X29))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(X29,X30))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_2])])])).
cnf(c_0_112, negated_conjecture, (v1_partfun1(esk3_0,esk1_0)|v1_partfun1(esk3_0,X1)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_108, c_0_109])).
cnf(c_0_113, negated_conjecture, (v1_partfun1(esk3_0,esk1_0)|v1_xboole_0(k6_partfun1(esk3_0))), inference(spm,[status(thm)],[c_0_110, c_0_76])).
cnf(c_0_114, negated_conjecture, (~v1_funct_1(esk3_0)|~v1_funct_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_115, plain, (v1_funct_2(X1,X2,X3)|~v1_funct_1(X1)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_116, negated_conjecture, (v1_partfun1(esk3_0,k6_partfun1(esk3_0))|v1_partfun1(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_112, c_0_113])).
cnf(c_0_117, negated_conjecture, (k6_partfun1(esk3_0)=esk1_0|v1_partfun1(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_96, c_0_107])).
cnf(c_0_118, plain, (r1_tarski(k2_relat_1(esk8_4(X1,X2,X3,X4)),X2)|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_119, negated_conjecture, (~r1_tarski(esk3_0,k2_zfmisc_1(esk1_0,esk2_0))|~v1_funct_2(esk3_0,esk1_0,esk2_0)|~v1_funct_1(esk3_0)), inference(spm,[status(thm)],[c_0_114, c_0_26])).
cnf(c_0_120, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_115, c_0_26])).
cnf(c_0_121, negated_conjecture, (v1_partfun1(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_122, plain, (r1_tarski(k2_relat_1(esk8_4(X1,X2,k1_funct_2(X1,X2),X3)),X2)|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_118])).
cnf(c_0_123, negated_conjecture, (~r1_tarski(esk3_0,k2_zfmisc_1(esk1_0,esk2_0))|~v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_119, c_0_55])])).
cnf(c_0_124, negated_conjecture, (v1_funct_2(esk3_0,esk1_0,X1)|~r1_tarski(esk3_0,k2_zfmisc_1(esk1_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121]), c_0_55])])).
cnf(c_0_125, plain, (r1_tarski(k2_relat_1(X1),X2)|~r2_hidden(X1,k1_funct_2(X3,X2))), inference(spm,[status(thm)],[c_0_122, c_0_30])).
cnf(c_0_126, negated_conjecture, (~r1_tarski(esk3_0,k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_123, c_0_124])).
cnf(c_0_127, plain, (r1_tarski(X1,k2_zfmisc_1(X2,X3))|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X3)), inference(spm,[status(thm)],[c_0_94, c_0_77])).
cnf(c_0_128, negated_conjecture, (r1_tarski(k2_relat_1(esk3_0),esk2_0)), inference(spm,[status(thm)],[c_0_125, c_0_41])).
cnf(c_0_129, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_127]), c_0_54]), c_0_53]), c_0_37]), c_0_128])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 130
# Proof object clause steps            : 97
# Proof object formula steps           : 33
# Proof object conjectures             : 39
# Proof object clause conjectures      : 36
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 63
# Proof object simplifying inferences  : 37
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 45
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 42
# Processed clauses                    : 4877
# ...of these trivial                  : 12
# ...subsumed                          : 3530
# ...remaining for further processing  : 1335
# Other redundant clauses eliminated   : 15
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 170
# Backward-rewritten                   : 220
# Generated clauses                    : 39260
# ...of the previous two non-trivial   : 38025
# Contextual simplify-reflections      : 60
# Paramodulations                      : 39248
# Factorizations                       : 0
# Equation resolutions                 : 15
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 892
#    Positive orientable unit clauses  : 114
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 775
# Current number of unprocessed clauses: 31186
# ...number of literals in the above   : 111827
# Current number of archived formulas  : 0
# Current number of archived clauses   : 431
# Clause-clause subsumption calls (NU) : 329528
# Rec. Clause-clause subsumption calls : 229625
# Non-unit clause-clause subsumptions  : 3722
# Unit Clause-clause subsumption calls : 5865
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1161
# BW rewrite match successes           : 13
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 523119

# -------------------------------------------------
# User time                : 0.398 s
# System time              : 0.020 s
# Total time               : 0.417 s
# Maximum resident set size: 29756 pages
