# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/20ecf13df45e4be8a94179db0a798094.E_input', t4_xboole_0)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/20ecf13df45e4be8a94179db0a798094.E_input', t12_setfam_1)).
fof(l24_zfmisc_1, axiom, ![X1, X2]:~((r1_xboole_0(k1_tarski(X1),X2)&r2_hidden(X1,X2))), file('data/MPTP/proofs/20ecf13df45e4be8a94179db0a798094.E_input', l24_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/20ecf13df45e4be8a94179db0a798094.E_input', t69_enumset1)).
fof(t168_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k10_relat_1(X2,X1)=k10_relat_1(X2,k3_xboole_0(k2_relat_1(X2),X1))), file('data/MPTP/proofs/20ecf13df45e4be8a94179db0a798094.E_input', t168_relat_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/20ecf13df45e4be8a94179db0a798094.E_input', t7_xboole_0)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/20ecf13df45e4be8a94179db0a798094.E_input', t65_xboole_1)).
fof(t173_relat_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>(k10_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k2_relat_1(X2),X1))), file('data/MPTP/proofs/20ecf13df45e4be8a94179db0a798094.E_input', t173_relat_1)).
fof(t166_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k10_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k2_relat_1(X3))&(r2_hidden(k4_tarski(X1,X4),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/20ecf13df45e4be8a94179db0a798094.E_input', t166_relat_1)).
fof(d14_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(X3=k10_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(k4_tarski(X4,X5),X1)&r2_hidden(X5,X2))))), file('data/MPTP/proofs/20ecf13df45e4be8a94179db0a798094.E_input', d14_relat_1)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/20ecf13df45e4be8a94179db0a798094.E_input', d4_xboole_0)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/20ecf13df45e4be8a94179db0a798094.E_input', d5_relat_1)).
fof(c_0_12, plain, ![X8, X9, X11, X12, X13]:((r1_xboole_0(X8,X9)|r2_hidden(esk3_2(X8,X9),k3_xboole_0(X8,X9)))&(~r2_hidden(X13,k3_xboole_0(X11,X12))|~r1_xboole_0(X11,X12))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_13, plain, ![X24, X25]:k1_setfam_1(k2_tarski(X24,X25))=k3_xboole_0(X24,X25), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_14, plain, ![X31, X32]:(~r1_xboole_0(k1_tarski(X31),X32)|~r2_hidden(X31,X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l24_zfmisc_1])])).
fof(c_0_15, plain, ![X23]:k2_tarski(X23,X23)=k1_tarski(X23), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_16, plain, ![X45, X46]:(~v1_relat_1(X46)|k10_relat_1(X46,X45)=k10_relat_1(X46,k3_xboole_0(k2_relat_1(X46),X45))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t168_relat_1])])).
cnf(c_0_17, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_19, plain, ![X59]:(X59=k1_xboole_0|r2_hidden(esk12_1(X59),X59)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_20, plain, (~r1_xboole_0(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_22, plain, ![X33]:r1_xboole_0(X33,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_23, plain, (k10_relat_1(X1,X2)=k10_relat_1(X1,k3_xboole_0(k2_relat_1(X1),X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (~r1_xboole_0(X2,X3)|~r2_hidden(X1,k1_setfam_1(k2_tarski(X2,X3)))), inference(rw,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_25, plain, (X1=k1_xboole_0|r2_hidden(esk12_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_26, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>(k10_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k2_relat_1(X2),X1)))), inference(assume_negation,[status(cth)],[t173_relat_1])).
cnf(c_0_27, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(k2_tarski(X1,X1),X2)), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_28, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_29, plain, ![X26, X27, X28, X30]:(((r2_hidden(esk5_3(X26,X27,X28),k2_relat_1(X28))|~r2_hidden(X26,k10_relat_1(X28,X27))|~v1_relat_1(X28))&((r2_hidden(k4_tarski(X26,esk5_3(X26,X27,X28)),X28)|~r2_hidden(X26,k10_relat_1(X28,X27))|~v1_relat_1(X28))&(r2_hidden(esk5_3(X26,X27,X28),X27)|~r2_hidden(X26,k10_relat_1(X28,X27))|~v1_relat_1(X28))))&(~r2_hidden(X30,k2_relat_1(X28))|(~r2_hidden(k4_tarski(X26,X30),X28)|~r2_hidden(X30,X27))|r2_hidden(X26,k10_relat_1(X28,X27))|~v1_relat_1(X28))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t166_relat_1])])])])])).
cnf(c_0_30, plain, (k10_relat_1(X1,X2)=k10_relat_1(X1,k1_setfam_1(k2_tarski(k2_relat_1(X1),X2)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_23, c_0_18])).
cnf(c_0_31, plain, (k1_setfam_1(k2_tarski(X1,X2))=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
fof(c_0_32, negated_conjecture, (v1_relat_1(esk2_0)&((k10_relat_1(esk2_0,esk1_0)!=k1_xboole_0|~r1_xboole_0(k2_relat_1(esk2_0),esk1_0))&(k10_relat_1(esk2_0,esk1_0)=k1_xboole_0|r1_xboole_0(k2_relat_1(esk2_0),esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])).
cnf(c_0_33, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34, plain, (r2_hidden(esk5_3(X1,X2,X3),X2)|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_35, plain, (k10_relat_1(X1,k1_xboole_0)=k10_relat_1(X1,X2)|~r1_xboole_0(k2_relat_1(X1),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_36, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)=k1_xboole_0|r1_xboole_0(k2_relat_1(esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_37, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, plain, (~r2_hidden(X1,k10_relat_1(X2,k1_xboole_0))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
fof(c_0_39, plain, ![X47, X48, X49, X50, X52, X53, X54, X55, X57]:((((r2_hidden(k4_tarski(X50,esk9_4(X47,X48,X49,X50)),X47)|~r2_hidden(X50,X49)|X49!=k10_relat_1(X47,X48)|~v1_relat_1(X47))&(r2_hidden(esk9_4(X47,X48,X49,X50),X48)|~r2_hidden(X50,X49)|X49!=k10_relat_1(X47,X48)|~v1_relat_1(X47)))&(~r2_hidden(k4_tarski(X52,X53),X47)|~r2_hidden(X53,X48)|r2_hidden(X52,X49)|X49!=k10_relat_1(X47,X48)|~v1_relat_1(X47)))&((~r2_hidden(esk10_3(X47,X54,X55),X55)|(~r2_hidden(k4_tarski(esk10_3(X47,X54,X55),X57),X47)|~r2_hidden(X57,X54))|X55=k10_relat_1(X47,X54)|~v1_relat_1(X47))&((r2_hidden(k4_tarski(esk10_3(X47,X54,X55),esk11_3(X47,X54,X55)),X47)|r2_hidden(esk10_3(X47,X54,X55),X55)|X55=k10_relat_1(X47,X54)|~v1_relat_1(X47))&(r2_hidden(esk11_3(X47,X54,X55),X54)|r2_hidden(esk10_3(X47,X54,X55),X55)|X55=k10_relat_1(X47,X54)|~v1_relat_1(X47))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d14_relat_1])])])])])])).
cnf(c_0_40, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)=k10_relat_1(esk2_0,k1_xboole_0)|k10_relat_1(esk2_0,esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37])])).
cnf(c_0_41, plain, (k10_relat_1(X1,k1_xboole_0)=k1_xboole_0|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_38, c_0_25])).
cnf(c_0_42, plain, (r2_hidden(esk9_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k10_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_43, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)=k1_xboole_0|k10_relat_1(esk2_0,k1_xboole_0)!=k1_xboole_0), inference(ef,[status(thm)],[c_0_40])).
cnf(c_0_44, negated_conjecture, (k10_relat_1(esk2_0,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_41, c_0_37])).
cnf(c_0_45, plain, (X1!=k10_relat_1(X2,k1_xboole_0)|~r2_hidden(X3,X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_33, c_0_42])).
fof(c_0_46, plain, ![X14, X15, X16, X17, X18, X19, X20, X21]:((((r2_hidden(X17,X14)|~r2_hidden(X17,X16)|X16!=k3_xboole_0(X14,X15))&(r2_hidden(X17,X15)|~r2_hidden(X17,X16)|X16!=k3_xboole_0(X14,X15)))&(~r2_hidden(X18,X14)|~r2_hidden(X18,X15)|r2_hidden(X18,X16)|X16!=k3_xboole_0(X14,X15)))&((~r2_hidden(esk4_3(X19,X20,X21),X21)|(~r2_hidden(esk4_3(X19,X20,X21),X19)|~r2_hidden(esk4_3(X19,X20,X21),X20))|X21=k3_xboole_0(X19,X20))&((r2_hidden(esk4_3(X19,X20,X21),X19)|r2_hidden(esk4_3(X19,X20,X21),X21)|X21=k3_xboole_0(X19,X20))&(r2_hidden(esk4_3(X19,X20,X21),X20)|r2_hidden(esk4_3(X19,X20,X21),X21)|X21=k3_xboole_0(X19,X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_47, plain, (r2_hidden(X1,X5)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X2,X4)|X5!=k10_relat_1(X3,X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_48, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_44])])).
cnf(c_0_49, negated_conjecture, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_44]), c_0_37])])).
cnf(c_0_50, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_51, plain, ![X34, X35, X36, X38, X39, X40, X41, X43]:(((~r2_hidden(X36,X35)|r2_hidden(k4_tarski(esk6_3(X34,X35,X36),X36),X34)|X35!=k2_relat_1(X34))&(~r2_hidden(k4_tarski(X39,X38),X34)|r2_hidden(X38,X35)|X35!=k2_relat_1(X34)))&((~r2_hidden(esk7_2(X40,X41),X41)|~r2_hidden(k4_tarski(X43,esk7_2(X40,X41)),X40)|X41=k2_relat_1(X40))&(r2_hidden(esk7_2(X40,X41),X41)|r2_hidden(k4_tarski(esk8_2(X40,X41),esk7_2(X40,X41)),X40)|X41=k2_relat_1(X40)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_52, negated_conjecture, (X1!=k1_xboole_0|~r2_hidden(k4_tarski(X2,X3),esk2_0)|~r2_hidden(X3,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_37])]), c_0_49])).
cnf(c_0_53, plain, (r2_hidden(X1,X2)|X3!=k1_setfam_1(k2_tarski(X4,X2))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_50, c_0_18])).
cnf(c_0_54, plain, (r2_hidden(X2,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_55, negated_conjecture, (~r2_hidden(k4_tarski(X1,X2),esk2_0)|~r2_hidden(X2,esk1_0)), inference(er,[status(thm)],[c_0_52])).
cnf(c_0_56, plain, (r2_hidden(k4_tarski(X1,esk5_3(X1,X2,X3)),X3)|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_57, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_setfam_1(k2_tarski(X3,X2)))), inference(er,[status(thm)],[c_0_53])).
cnf(c_0_58, plain, (r2_hidden(X3,k10_relat_1(X2,X4))|~r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)|~r2_hidden(X1,X4)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_59, plain, (r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)), inference(er,[status(thm)],[c_0_54])).
cnf(c_0_60, negated_conjecture, (~r2_hidden(esk5_3(X1,X2,esk2_0),esk1_0)|~r2_hidden(X1,k10_relat_1(esk2_0,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_37])])).
cnf(c_0_61, plain, (r2_hidden(esk5_3(X1,k1_setfam_1(k2_tarski(X2,X3)),X4),X3)|~r2_hidden(X1,k10_relat_1(X4,k1_setfam_1(k2_tarski(X2,X3))))|~v1_relat_1(X4)), inference(spm,[status(thm)],[c_0_57, c_0_34])).
cnf(c_0_62, plain, (r2_hidden(X1,k10_relat_1(X2,X3))|~r2_hidden(k4_tarski(X1,X4),X2)|~r2_hidden(X4,X3)|~v1_relat_1(X2)), inference(csr,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_63, plain, (r2_hidden(k4_tarski(esk6_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_64, negated_conjecture, (~r2_hidden(X1,k10_relat_1(esk2_0,k1_setfam_1(k2_tarski(X2,esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_37])])).
cnf(c_0_65, plain, (r2_hidden(esk6_3(X1,X2,X3),k10_relat_1(X1,X4))|X2!=k2_relat_1(X1)|~r2_hidden(X3,X4)|~r2_hidden(X3,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_66, negated_conjecture, (k10_relat_1(esk2_0,k1_setfam_1(k2_tarski(X1,esk1_0)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_64, c_0_25])).
cnf(c_0_67, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_68, negated_conjecture, (X1!=k2_relat_1(esk2_0)|~r2_hidden(X2,k1_setfam_1(k2_tarski(X3,esk1_0)))|~r2_hidden(X2,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_37])]), c_0_33])).
cnf(c_0_69, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk3_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_70, plain, (r2_hidden(X1,X2)|X3!=k1_setfam_1(k2_tarski(X2,X4))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_67, c_0_18])).
cnf(c_0_71, negated_conjecture, (~r2_hidden(X1,k1_setfam_1(k2_tarski(X2,esk1_0)))|~r2_hidden(X1,k2_relat_1(esk2_0))), inference(er,[status(thm)],[c_0_68])).
cnf(c_0_72, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk3_2(X1,X2),k1_setfam_1(k2_tarski(X1,X2)))), inference(rw,[status(thm)],[c_0_69, c_0_18])).
cnf(c_0_73, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_setfam_1(k2_tarski(X2,X3)))), inference(er,[status(thm)],[c_0_70])).
cnf(c_0_74, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)!=k1_xboole_0|~r1_xboole_0(k2_relat_1(esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_75, negated_conjecture, (r1_xboole_0(X1,esk1_0)|~r2_hidden(esk3_2(X1,esk1_0),k2_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_76, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_73, c_0_72])).
cnf(c_0_77, negated_conjecture, (~r1_xboole_0(k2_relat_1(esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_48])])).
cnf(c_0_78, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 79
# Proof object clause steps            : 54
# Proof object formula steps           : 25
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 26
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 29
# Processed clauses                    : 3438
# ...of these trivial                  : 17
# ...subsumed                          : 2871
# ...remaining for further processing  : 550
# Other redundant clauses eliminated   : 333
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 38
# Backward-rewritten                   : 16
# Generated clauses                    : 22764
# ...of the previous two non-trivial   : 20791
# Contextual simplify-reflections      : 9
# Paramodulations                      : 22218
# Factorizations                       : 69
# Equation resolutions                 : 477
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 467
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 451
# Current number of unprocessed clauses: 17263
# ...number of literals in the above   : 92814
# Current number of archived formulas  : 0
# Current number of archived clauses   : 85
# Clause-clause subsumption calls (NU) : 67975
# Rec. Clause-clause subsumption calls : 31357
# Non-unit clause-clause subsumptions  : 2158
# Unit Clause-clause subsumption calls : 198
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 10
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 360519

# -------------------------------------------------
# User time                : 0.296 s
# System time              : 0.000 s
# Total time               : 0.296 s
# Maximum resident set size: 28792 pages
