# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', t70_enumset1)).
fof(t95_xboole_1, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2)), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', t95_xboole_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', t72_enumset1)).
fof(t43_zfmisc_1, conjecture, ![X1, X2, X3]:~((k1_tarski(X1)=k2_xboole_0(X2,X3)&(~((X2=k1_tarski(X1)&X3=k1_tarski(X1)))&(~((X2=k1_xboole_0&X3=k1_tarski(X1)))&~((X2=k1_tarski(X1)&X3=k1_xboole_0)))))), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', t43_zfmisc_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', idempotence_k3_xboole_0)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', t75_enumset1)).
fof(idempotence_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', idempotence_k2_xboole_0)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', t7_xboole_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', t69_enumset1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', t36_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', t100_xboole_1)).
fof(t92_xboole_1, axiom, ![X1]:k5_xboole_0(X1,X1)=k1_xboole_0, file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', t92_xboole_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', d1_tarski)).
fof(t91_xboole_1, axiom, ![X1, X2, X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3)), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', t91_xboole_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', d3_tarski)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', d10_xboole_0)).
fof(l1_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', l1_zfmisc_1)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', t1_xboole_1)).
fof(t1_xboole_0, axiom, ![X1, X2, X3]:(r2_hidden(X1,k5_xboole_0(X2,X3))<=>~((r2_hidden(X1,X2)<=>r2_hidden(X1,X3)))), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', t1_xboole_0)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/215bb785a5eb41939fdf8fa63ffeadd0.E_input', t2_xboole_1)).
fof(c_0_24, plain, ![X8, X9]:k3_tarski(k2_tarski(X8,X9))=k2_xboole_0(X8,X9), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_25, plain, ![X10, X11]:k1_enumset1(X10,X10,X11)=k2_tarski(X10,X11), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_26, plain, ![X12, X13]:k3_xboole_0(X12,X13)=k5_xboole_0(k5_xboole_0(X12,X13),k2_xboole_0(X12,X13)), inference(variable_rename,[status(thm)],[t95_xboole_1])).
cnf(c_0_27, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_28, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_29, plain, ![X14, X15, X16]:k2_enumset1(X14,X14,X15,X16)=k1_enumset1(X14,X15,X16), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_30, plain, ![X17, X18, X19, X20]:k3_enumset1(X17,X17,X18,X19,X20)=k2_enumset1(X17,X18,X19,X20), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_31, negated_conjecture, ~(![X1, X2, X3]:~((k1_tarski(X1)=k2_xboole_0(X2,X3)&(~((X2=k1_tarski(X1)&X3=k1_tarski(X1)))&(~((X2=k1_xboole_0&X3=k1_tarski(X1)))&~((X2=k1_tarski(X1)&X3=k1_xboole_0))))))), inference(assume_negation,[status(cth)],[t43_zfmisc_1])).
fof(c_0_32, plain, ![X24]:k3_xboole_0(X24,X24)=X24, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_33, plain, (k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_35, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_37, plain, ![X25, X26, X27, X28, X29]:k4_enumset1(X25,X25,X26,X27,X28,X29)=k3_enumset1(X25,X26,X27,X28,X29), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_38, plain, ![X30, X31, X32, X33, X34, X35]:k5_enumset1(X30,X30,X31,X32,X33,X34,X35)=k4_enumset1(X30,X31,X32,X33,X34,X35), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_39, plain, ![X36, X37, X38, X39, X40, X41, X42]:k6_enumset1(X36,X36,X37,X38,X39,X40,X41,X42)=k5_enumset1(X36,X37,X38,X39,X40,X41,X42), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_40, plain, ![X43]:k2_xboole_0(X43,X43)=X43, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k2_xboole_0])])).
fof(c_0_41, plain, ![X44, X45]:r1_tarski(X44,k2_xboole_0(X44,X45)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_42, negated_conjecture, (k1_tarski(esk1_0)=k2_xboole_0(esk2_0,esk3_0)&((esk2_0!=k1_tarski(esk1_0)|esk3_0!=k1_tarski(esk1_0))&((esk2_0!=k1_xboole_0|esk3_0!=k1_tarski(esk1_0))&(esk2_0!=k1_tarski(esk1_0)|esk3_0!=k1_xboole_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])])).
fof(c_0_43, plain, ![X46]:k2_tarski(X46,X46)=k1_tarski(X46), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_44, plain, ![X47, X48]:r1_tarski(k4_xboole_0(X47,X48),X47), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_45, plain, ![X49, X50]:k4_xboole_0(X49,X50)=k5_xboole_0(X49,k3_xboole_0(X49,X50)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_46, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_47, plain, (k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k3_enumset1(X1,X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_34]), c_0_35]), c_0_36])).
cnf(c_0_48, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_49, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_51, plain, ![X51]:k5_xboole_0(X51,X51)=k1_xboole_0, inference(variable_rename,[status(thm)],[t92_xboole_1])).
cnf(c_0_52, plain, (k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_53, plain, ![X52, X53, X54, X55, X56, X57]:(((~r2_hidden(X54,X53)|X54=X52|X53!=k1_tarski(X52))&(X55!=X52|r2_hidden(X55,X53)|X53!=k1_tarski(X52)))&((~r2_hidden(esk4_2(X56,X57),X57)|esk4_2(X56,X57)!=X56|X57=k1_tarski(X56))&(r2_hidden(esk4_2(X56,X57),X57)|esk4_2(X56,X57)=X56|X57=k1_tarski(X56)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_54, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_55, negated_conjecture, (k1_tarski(esk1_0)=k2_xboole_0(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_56, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_57, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_58, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_59, plain, ![X59, X60, X61]:k5_xboole_0(k5_xboole_0(X59,X60),X61)=k5_xboole_0(X59,k5_xboole_0(X60,X61)), inference(variable_rename,[status(thm)],[t91_xboole_1])).
cnf(c_0_60, plain, (k5_xboole_0(k5_xboole_0(X1,X1),k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_47]), c_0_48]), c_0_49]), c_0_50])).
cnf(c_0_61, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_62, plain, (k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_34]), c_0_35]), c_0_36]), c_0_48]), c_0_49]), c_0_50])).
cnf(c_0_63, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_53])).
fof(c_0_64, plain, ![X62, X63, X64, X65, X66]:((~r1_tarski(X62,X63)|(~r2_hidden(X64,X62)|r2_hidden(X64,X63)))&((r2_hidden(esk5_2(X65,X66),X65)|r1_tarski(X65,X66))&(~r2_hidden(esk5_2(X65,X66),X66)|r1_tarski(X65,X66)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_65, plain, (r1_tarski(X1,k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_34]), c_0_35]), c_0_36]), c_0_48]), c_0_49]), c_0_50])).
cnf(c_0_66, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=k3_tarski(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_56]), c_0_28]), c_0_34]), c_0_35]), c_0_35]), c_0_36]), c_0_36]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_50]), c_0_50])).
cnf(c_0_67, plain, (r1_tarski(k5_xboole_0(X1,k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2)))),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_58]), c_0_47]), c_0_48]), c_0_49]), c_0_50])).
cnf(c_0_68, plain, (k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_69, plain, (k5_xboole_0(k1_xboole_0,X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_61]), c_0_62])).
cnf(c_0_70, plain, (X1=X3|X2!=k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_56]), c_0_28]), c_0_35]), c_0_36]), c_0_48]), c_0_49]), c_0_50])).
cnf(c_0_71, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_72, negated_conjecture, (r1_tarski(esk2_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_73, plain, (r1_tarski(k5_xboole_0(X1,k5_xboole_0(X1,k5_xboole_0(X2,k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))))),X1)), inference(rw,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_74, plain, (k5_xboole_0(X1,k5_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_61]), c_0_69])).
fof(c_0_75, plain, ![X68, X69]:(((r1_tarski(X68,X69)|X68!=X69)&(r1_tarski(X69,X68)|X68!=X69))&(~r1_tarski(X68,X69)|~r1_tarski(X69,X68)|X68=X69)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_76, plain, ![X70, X71]:((~r1_tarski(k1_tarski(X70),X71)|r2_hidden(X70,X71))&(~r2_hidden(X70,X71)|r1_tarski(k1_tarski(X70),X71))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])).
cnf(c_0_77, plain, (X1=X2|~r2_hidden(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2))), inference(er,[status(thm)],[c_0_70])).
cnf(c_0_78, negated_conjecture, (r2_hidden(X1,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_79, plain, (r1_tarski(k5_xboole_0(X1,k3_tarski(k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X1))),X2)), inference(rw,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_80, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_81, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_82, plain, (k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X1,X2)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_61, c_0_68])).
cnf(c_0_83, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(spm,[status(thm)],[c_0_74, c_0_61])).
fof(c_0_84, plain, ![X72, X73, X74]:(~r1_tarski(X72,X73)|~r1_tarski(X73,X74)|r1_tarski(X72,X74)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
cnf(c_0_85, negated_conjecture, (X1=esk1_0|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_86, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_87, negated_conjecture, (r1_tarski(k5_xboole_0(esk3_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)),esk2_0)), inference(spm,[status(thm)],[c_0_79, c_0_66])).
cnf(c_0_88, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=esk2_0|~r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),esk2_0)), inference(spm,[status(thm)],[c_0_80, c_0_72])).
cnf(c_0_89, plain, (r1_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_56]), c_0_28]), c_0_35]), c_0_36]), c_0_48]), c_0_49]), c_0_50])).
cnf(c_0_90, plain, (k5_xboole_0(X1,k5_xboole_0(X2,X1))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_82]), c_0_83])).
cnf(c_0_91, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_92, negated_conjecture, (esk5_2(esk2_0,X1)=esk1_0|r1_tarski(esk2_0,X1)), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_93, negated_conjecture, (k5_xboole_0(esk3_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))=esk2_0|~r1_tarski(esk2_0,k5_xboole_0(esk3_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))), inference(spm,[status(thm)],[c_0_80, c_0_87])).
cnf(c_0_94, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=esk2_0|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_95, plain, (k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_74, c_0_90])).
cnf(c_0_96, plain, (r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_97, negated_conjecture, (r1_tarski(X1,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))|~r1_tarski(X1,esk2_0)), inference(spm,[status(thm)],[c_0_91, c_0_72])).
cnf(c_0_98, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|r1_tarski(esk2_0,X1)), inference(spm,[status(thm)],[c_0_86, c_0_92])).
cnf(c_0_99, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_100, negated_conjecture, (k5_xboole_0(esk2_0,esk3_0)=esk2_0|~r2_hidden(esk1_0,esk2_0)|~r1_tarski(esk2_0,k5_xboole_0(esk2_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_95]), c_0_95])).
cnf(c_0_101, negated_conjecture, (r1_tarski(esk2_0,X1)|~r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_96, c_0_92])).
fof(c_0_102, plain, ![X75, X76, X77]:(((~r2_hidden(X75,X76)|~r2_hidden(X75,X77)|~r2_hidden(X75,k5_xboole_0(X76,X77)))&(r2_hidden(X75,X76)|r2_hidden(X75,X77)|~r2_hidden(X75,k5_xboole_0(X76,X77))))&((~r2_hidden(X75,X76)|r2_hidden(X75,X77)|r2_hidden(X75,k5_xboole_0(X76,X77)))&(~r2_hidden(X75,X77)|r2_hidden(X75,X76)|r2_hidden(X75,k5_xboole_0(X76,X77))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_0])])])).
cnf(c_0_103, negated_conjecture, (esk2_0!=k1_tarski(esk1_0)|esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_104, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=X1|~r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),X1)|~r1_tarski(X1,esk2_0)), inference(spm,[status(thm)],[c_0_80, c_0_97])).
cnf(c_0_105, negated_conjecture, (k5_xboole_0(esk3_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))=esk2_0|r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_93, c_0_98])).
cnf(c_0_106, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_99, c_0_56]), c_0_28]), c_0_35]), c_0_36]), c_0_48]), c_0_49]), c_0_50])).
cnf(c_0_107, negated_conjecture, (k5_xboole_0(esk2_0,esk3_0)=esk2_0|~r2_hidden(esk1_0,k5_xboole_0(esk2_0,esk3_0))|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_108, plain, (r2_hidden(X1,X3)|r2_hidden(X1,k5_xboole_0(X2,X3))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_109, negated_conjecture, (esk3_0!=k1_xboole_0|esk2_0!=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_56]), c_0_28]), c_0_35]), c_0_36]), c_0_48]), c_0_49]), c_0_50])).
cnf(c_0_110, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=X1|~r2_hidden(esk1_0,X1)|~r1_tarski(X1,esk2_0)), inference(spm,[status(thm)],[c_0_104, c_0_89])).
cnf(c_0_111, negated_conjecture, (r1_tarski(k5_xboole_0(esk2_0,esk3_0),esk2_0)|~r2_hidden(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_94]), c_0_95])).
cnf(c_0_112, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=k5_xboole_0(esk2_0,esk3_0)|r2_hidden(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_105]), c_0_95])).
cnf(c_0_113, plain, (r2_hidden(X1,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_106])])).
cnf(c_0_114, negated_conjecture, (k5_xboole_0(esk2_0,esk3_0)=esk2_0|r2_hidden(esk1_0,esk3_0)|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_115, negated_conjecture, (esk3_0!=k1_xboole_0|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_109, c_0_94])).
cnf(c_0_116, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=k5_xboole_0(esk2_0,esk3_0)|~r2_hidden(esk1_0,k5_xboole_0(esk2_0,esk3_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_112])).
cnf(c_0_117, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_118, negated_conjecture, (r2_hidden(esk1_0,k5_xboole_0(esk2_0,esk3_0))|r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_113, c_0_112])).
cnf(c_0_119, negated_conjecture, (r2_hidden(esk1_0,esk3_0)|~r2_hidden(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_114]), c_0_61]), c_0_115])).
cnf(c_0_120, negated_conjecture, (X1=esk1_0|~r2_hidden(esk1_0,k5_xboole_0(esk2_0,esk3_0))|~r2_hidden(X1,k5_xboole_0(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_77, c_0_116])).
cnf(c_0_121, plain, (r2_hidden(X1,X3)|r2_hidden(X1,k5_xboole_0(X3,X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_122, negated_conjecture, (r2_hidden(esk1_0,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_119])).
cnf(c_0_123, negated_conjecture, (X1=esk1_0|r2_hidden(esk1_0,esk2_0)|~r2_hidden(X1,k5_xboole_0(esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121]), c_0_122])])).
fof(c_0_124, plain, ![X78]:r1_tarski(k1_xboole_0,X78), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_125, negated_conjecture, (X1=esk1_0|r2_hidden(esk1_0,esk2_0)|~r2_hidden(X1,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_121]), c_0_85])).
cnf(c_0_126, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_127, negated_conjecture, (esk5_2(esk3_0,X1)=esk1_0|r2_hidden(esk1_0,esk2_0)|r1_tarski(esk3_0,X1)), inference(spm,[status(thm)],[c_0_125, c_0_86])).
cnf(c_0_128, negated_conjecture, (esk2_0!=k1_xboole_0|esk3_0!=k1_tarski(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_129, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_80, c_0_126])).
cnf(c_0_130, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|r1_tarski(esk3_0,X1)|~r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_96, c_0_127])).
cnf(c_0_131, negated_conjecture, (esk5_2(k5_xboole_0(esk2_0,esk3_0),X1)=esk1_0|r2_hidden(esk1_0,esk2_0)|r1_tarski(k5_xboole_0(esk2_0,esk3_0),X1)), inference(spm,[status(thm)],[c_0_123, c_0_86])).
cnf(c_0_132, negated_conjecture, (esk2_0!=k1_xboole_0|esk3_0!=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_128, c_0_56]), c_0_28]), c_0_35]), c_0_36]), c_0_48]), c_0_49]), c_0_50])).
cnf(c_0_133, negated_conjecture, (k1_xboole_0=esk2_0|r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_129, c_0_98])).
cnf(c_0_134, negated_conjecture, (X1=esk3_0|r2_hidden(esk1_0,esk2_0)|~r2_hidden(esk1_0,X1)|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_80, c_0_130])).
cnf(c_0_135, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|r1_tarski(k5_xboole_0(esk2_0,esk3_0),X1)|~r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_96, c_0_131])).
cnf(c_0_136, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|k5_xboole_0(esk2_0,esk3_0)!=esk3_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_112]), c_0_133])).
cnf(c_0_137, negated_conjecture, (r2_hidden(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_122])]), c_0_118]), c_0_136])).
cnf(c_0_138, negated_conjecture, (r1_tarski(k5_xboole_0(esk2_0,esk3_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_111, c_0_137])])).
cnf(c_0_139, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,k5_xboole_0(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_71, c_0_138])).
cnf(c_0_140, negated_conjecture, (esk2_0!=k1_tarski(esk1_0)|esk3_0!=k1_tarski(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_141, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_139, c_0_121])).
cnf(c_0_142, negated_conjecture, (esk2_0!=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)|esk3_0!=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_140, c_0_56]), c_0_56]), c_0_28]), c_0_28]), c_0_35]), c_0_35]), c_0_36]), c_0_36]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_50]), c_0_50])).
cnf(c_0_143, negated_conjecture, (r2_hidden(esk5_2(esk3_0,X1),esk2_0)|r1_tarski(esk3_0,X1)), inference(spm,[status(thm)],[c_0_141, c_0_86])).
cnf(c_0_144, negated_conjecture, (esk3_0!=esk2_0|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_142, c_0_94])).
cnf(c_0_145, negated_conjecture, (X1=esk2_0|~r2_hidden(esk1_0,X1)|~r1_tarski(X1,esk2_0)), inference(spm,[status(thm)],[c_0_80, c_0_101])).
cnf(c_0_146, negated_conjecture, (r1_tarski(esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_96, c_0_143])).
cnf(c_0_147, negated_conjecture, (esk3_0!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_144, c_0_137])])).
cnf(c_0_148, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146]), c_0_122])]), c_0_147]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 149
# Proof object clause steps            : 100
# Proof object formula steps           : 49
# Proof object conjectures             : 55
# Proof object clause conjectures      : 52
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 32
# Proof object initial formulas used   : 24
# Proof object generating inferences   : 48
# Proof object simplifying inferences  : 119
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 24
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 38
# Removed in clause preprocessing      : 10
# Initial clauses in saturation        : 28
# Processed clauses                    : 453
# ...of these trivial                  : 10
# ...subsumed                          : 256
# ...remaining for further processing  : 187
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 64
# Generated clauses                    : 1505
# ...of the previous two non-trivial   : 1217
# Contextual simplify-reflections      : 11
# Paramodulations                      : 1500
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 88
#    Positive orientable unit clauses  : 25
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 57
# Current number of unprocessed clauses: 782
# ...number of literals in the above   : 2263
# Current number of archived formulas  : 0
# Current number of archived clauses   : 105
# Clause-clause subsumption calls (NU) : 2256
# Rec. Clause-clause subsumption calls : 1633
# Non-unit clause-clause subsumptions  : 209
# Unit Clause-clause subsumption calls : 123
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 77
# BW rewrite match successes           : 61
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 19537

# -------------------------------------------------
# User time                : 0.022 s
# System time              : 0.000 s
# Total time               : 0.022 s
# Maximum resident set size: 26384 pages
