# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t39_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2))), file('data/MPTP/proofs/21856b984aa14422982ce2c4a348ddf0.E_input', t39_yellow_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/21856b984aa14422982ce2c4a348ddf0.E_input', t24_orders_2)).
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/21856b984aa14422982ce2c4a348ddf0.E_input', d9_yellow_0)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/21856b984aa14422982ce2c4a348ddf0.E_input', dt_k1_yellow_0)).
fof(t18_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))&((r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))=>(X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3))))))))), file('data/MPTP/proofs/21856b984aa14422982ce2c4a348ddf0.E_input', t18_yellow_0)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/21856b984aa14422982ce2c4a348ddf0.E_input', redefinition_k6_domain_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/21856b984aa14422982ce2c4a348ddf0.E_input', t69_enumset1)).
fof(t7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2)))))))), file('data/MPTP/proofs/21856b984aa14422982ce2c4a348ddf0.E_input', t7_yellow_0)).
fof(t8_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4)))&(((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4))=>r1_lattice3(X1,k2_tarski(X3,X4),X2))&((r2_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2)))&((r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2))=>r2_lattice3(X1,k2_tarski(X3,X4),X2))))))))), file('data/MPTP/proofs/21856b984aa14422982ce2c4a348ddf0.E_input', t8_yellow_0)).
fof(dt_k2_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/21856b984aa14422982ce2c4a348ddf0.E_input', dt_k2_yellow_0)).
fof(d10_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_yellow_0(X1,X2)=>(X3=k2_yellow_0(X1,X2)<=>(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))))))), file('data/MPTP/proofs/21856b984aa14422982ce2c4a348ddf0.E_input', d10_yellow_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/21856b984aa14422982ce2c4a348ddf0.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/21856b984aa14422982ce2c4a348ddf0.E_input', dt_l1_orders_2)).
fof(t19_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))))), file('data/MPTP/proofs/21856b984aa14422982ce2c4a348ddf0.E_input', t19_yellow_0)).
fof(c_0_14, plain, ![X3, X2, X1]:(epred3_3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))&((r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))=>(X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3))))))), introduced(definition)).
fof(c_0_15, plain, ![X3, X2, X1]:(epred3_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))&((r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))=>(X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3))))))), inference(split_equiv,[status(thm)],[c_0_14])).
fof(c_0_16, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2)))), inference(assume_negation,[status(cth)],[t39_yellow_0])).
fof(c_0_17, plain, ![X52, X53, X54, X55, X56]:(((r1_orders_2(X54,X53,X55)|(X55!=k10_lattice3(X54,X53,X52)|~r1_yellow_0(X54,k2_tarski(X53,X52)))|~m1_subset_1(X55,u1_struct_0(X54))|~epred3_3(X54,X53,X52))&((r1_orders_2(X54,X52,X55)|(X55!=k10_lattice3(X54,X53,X52)|~r1_yellow_0(X54,k2_tarski(X53,X52)))|~m1_subset_1(X55,u1_struct_0(X54))|~epred3_3(X54,X53,X52))&(~m1_subset_1(X56,u1_struct_0(X54))|(~r1_orders_2(X54,X53,X56)|~r1_orders_2(X54,X52,X56)|r1_orders_2(X54,X55,X56))|(X55!=k10_lattice3(X54,X53,X52)|~r1_yellow_0(X54,k2_tarski(X53,X52)))|~m1_subset_1(X55,u1_struct_0(X54))|~epred3_3(X54,X53,X52))))&(((X55=k10_lattice3(X54,X53,X52)|(m1_subset_1(esk6_4(X52,X53,X54,X55),u1_struct_0(X54))|~r1_orders_2(X54,X52,X55)|~r1_orders_2(X54,X53,X55))|~m1_subset_1(X55,u1_struct_0(X54))|~epred3_3(X54,X53,X52))&(r1_yellow_0(X54,k2_tarski(X53,X52))|(m1_subset_1(esk6_4(X52,X53,X54,X55),u1_struct_0(X54))|~r1_orders_2(X54,X52,X55)|~r1_orders_2(X54,X53,X55))|~m1_subset_1(X55,u1_struct_0(X54))|~epred3_3(X54,X53,X52)))&((((X55=k10_lattice3(X54,X53,X52)|(r1_orders_2(X54,X53,esk6_4(X52,X53,X54,X55))|~r1_orders_2(X54,X52,X55)|~r1_orders_2(X54,X53,X55))|~m1_subset_1(X55,u1_struct_0(X54))|~epred3_3(X54,X53,X52))&(r1_yellow_0(X54,k2_tarski(X53,X52))|(r1_orders_2(X54,X53,esk6_4(X52,X53,X54,X55))|~r1_orders_2(X54,X52,X55)|~r1_orders_2(X54,X53,X55))|~m1_subset_1(X55,u1_struct_0(X54))|~epred3_3(X54,X53,X52)))&((X55=k10_lattice3(X54,X53,X52)|(r1_orders_2(X54,X52,esk6_4(X52,X53,X54,X55))|~r1_orders_2(X54,X52,X55)|~r1_orders_2(X54,X53,X55))|~m1_subset_1(X55,u1_struct_0(X54))|~epred3_3(X54,X53,X52))&(r1_yellow_0(X54,k2_tarski(X53,X52))|(r1_orders_2(X54,X52,esk6_4(X52,X53,X54,X55))|~r1_orders_2(X54,X52,X55)|~r1_orders_2(X54,X53,X55))|~m1_subset_1(X55,u1_struct_0(X54))|~epred3_3(X54,X53,X52))))&((X55=k10_lattice3(X54,X53,X52)|(~r1_orders_2(X54,X55,esk6_4(X52,X53,X54,X55))|~r1_orders_2(X54,X52,X55)|~r1_orders_2(X54,X53,X55))|~m1_subset_1(X55,u1_struct_0(X54))|~epred3_3(X54,X53,X52))&(r1_yellow_0(X54,k2_tarski(X53,X52))|(~r1_orders_2(X54,X55,esk6_4(X52,X53,X54,X55))|~r1_orders_2(X54,X52,X55)|~r1_orders_2(X54,X53,X55))|~m1_subset_1(X55,u1_struct_0(X54))|~epred3_3(X54,X53,X52)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])])])).
fof(c_0_18, plain, ![X32, X33]:(v2_struct_0(X32)|(~v3_orders_2(X32)|~l1_orders_2(X32))|(~m1_subset_1(X33,u1_struct_0(X32))|r1_orders_2(X32,X33,X33))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_19, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_16])])])])).
fof(c_0_20, plain, ![X23, X24, X25, X26]:(((r2_lattice3(X23,X24,X25)|X25!=k1_yellow_0(X23,X24)|~r1_yellow_0(X23,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))&(~m1_subset_1(X26,u1_struct_0(X23))|(~r2_lattice3(X23,X24,X26)|r1_orders_2(X23,X25,X26))|X25!=k1_yellow_0(X23,X24)|~r1_yellow_0(X23,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23)))&((m1_subset_1(esk4_3(X23,X24,X25),u1_struct_0(X23))|~r2_lattice3(X23,X24,X25)|X25=k1_yellow_0(X23,X24)|~r1_yellow_0(X23,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))&((r2_lattice3(X23,X24,esk4_3(X23,X24,X25))|~r2_lattice3(X23,X24,X25)|X25=k1_yellow_0(X23,X24)|~r1_yellow_0(X23,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))&(~r1_orders_2(X23,X25,esk4_3(X23,X24,X25))|~r2_lattice3(X23,X24,X25)|X25=k1_yellow_0(X23,X24)|~r1_yellow_0(X23,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
fof(c_0_21, plain, ![X19, X20]:(~l1_orders_2(X19)|m1_subset_1(k1_yellow_0(X19,X20),u1_struct_0(X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
cnf(c_0_22, plain, (X1=k10_lattice3(X2,X3,X4)|~r1_orders_2(X2,X1,esk6_4(X4,X3,X2,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~epred3_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (X1=k10_lattice3(X2,X3,X4)|r1_orders_2(X2,X3,esk6_4(X4,X3,X2,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~epred3_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (X1=k10_lattice3(X2,X3,X4)|r1_orders_2(X2,X4,esk6_4(X4,X3,X2,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~epred3_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_30, plain, (r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~r1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_31, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_32, plain, (r2_lattice3(X1,X2,X3)|X3!=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_33, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>epred3_3(X1,X2,X3)))), inference(apply_def,[status(thm)],[t18_yellow_0, c_0_14])).
cnf(c_0_34, plain, (k10_lattice3(X1,X2,X3)=X2|~epred3_3(X1,X2,X3)|~r1_orders_2(X1,X3,X2)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_35, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26]), c_0_27])]), c_0_28])).
cnf(c_0_36, plain, (k10_lattice3(X1,X2,X3)=X3|~epred3_3(X1,X2,X3)|~r1_orders_2(X1,X3,X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_22, c_0_29])).
cnf(c_0_37, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),X3)|~r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_30]), c_0_31])).
cnf(c_0_38, plain, (r2_lattice3(X1,X2,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_32]), c_0_31])).
fof(c_0_39, plain, ![X39, X40, X41]:(~v5_orders_2(X39)|~l1_orders_2(X39)|(~m1_subset_1(X40,u1_struct_0(X39))|(~m1_subset_1(X41,u1_struct_0(X39))|epred3_3(X39,X40,X41)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_33])])])).
fof(c_0_40, plain, ![X4, X3, X2, X1]:(epred1_4(X1,X2,X3,X4)<=>((r1_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4)))&(((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4))=>r1_lattice3(X1,k2_tarski(X3,X4),X2))&((r2_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2)))&((r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2))=>r2_lattice3(X1,k2_tarski(X3,X4),X2)))))), introduced(definition)).
fof(c_0_41, plain, ![X37, X38]:(v1_xboole_0(X37)|~m1_subset_1(X38,X37)|k6_domain_1(X37,X38)=k1_tarski(X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_42, plain, ![X18]:k2_tarski(X18,X18)=k1_tarski(X18), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_43, negated_conjecture, (k10_lattice3(esk1_0,esk2_0,X1)=esk2_0|~epred3_3(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_25])])).
cnf(c_0_44, plain, (k10_lattice3(X1,X2,k1_yellow_0(X1,X3))=k1_yellow_0(X1,X3)|~epred3_3(X1,X2,k1_yellow_0(X1,X3))|~r1_yellow_0(X1,X3)|~r1_orders_2(X1,X2,k1_yellow_0(X1,X3))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_31]), c_0_38])).
cnf(c_0_45, plain, (epred3_3(X1,X2,X3)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_46, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_47, plain, ![X4, X3, X2, X1]:(epred1_4(X1,X2,X3,X4)=>((r1_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4)))&(((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4))=>r1_lattice3(X1,k2_tarski(X3,X4),X2))&((r2_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2)))&((r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2))=>r2_lattice3(X1,k2_tarski(X3,X4),X2)))))), inference(split_equiv,[status(thm)],[c_0_40])).
fof(c_0_48, plain, ![X34, X35, X36]:((~r1_lattice3(X34,k1_tarski(X36),X35)|r1_orders_2(X34,X35,X36)|~m1_subset_1(X36,u1_struct_0(X34))|~m1_subset_1(X35,u1_struct_0(X34))|~l1_orders_2(X34))&((~r1_orders_2(X34,X35,X36)|r1_lattice3(X34,k1_tarski(X36),X35)|~m1_subset_1(X36,u1_struct_0(X34))|~m1_subset_1(X35,u1_struct_0(X34))|~l1_orders_2(X34))&((~r2_lattice3(X34,k1_tarski(X36),X35)|r1_orders_2(X34,X36,X35)|~m1_subset_1(X36,u1_struct_0(X34))|~m1_subset_1(X35,u1_struct_0(X34))|~l1_orders_2(X34))&(~r1_orders_2(X34,X36,X35)|r2_lattice3(X34,k1_tarski(X36),X35)|~m1_subset_1(X36,u1_struct_0(X34))|~m1_subset_1(X35,u1_struct_0(X34))|~l1_orders_2(X34))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])).
cnf(c_0_49, plain, (r1_yellow_0(X1,k2_tarski(X2,X3))|~r1_orders_2(X1,X4,esk6_4(X3,X2,X1,X4))|~r1_orders_2(X1,X3,X4)|~r1_orders_2(X1,X2,X4)|~m1_subset_1(X4,u1_struct_0(X1))|~epred3_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_50, plain, (r1_yellow_0(X1,k2_tarski(X2,X3))|r1_orders_2(X1,X3,esk6_4(X3,X2,X1,X4))|~r1_orders_2(X1,X3,X4)|~r1_orders_2(X1,X2,X4)|~m1_subset_1(X4,u1_struct_0(X1))|~epred3_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_51, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_52, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_53, negated_conjecture, (k1_yellow_0(esk1_0,X1)=esk2_0|~epred3_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),esk2_0)|~r1_orders_2(esk1_0,esk2_0,k1_yellow_0(esk1_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_26])])).
cnf(c_0_54, negated_conjecture, (epred3_3(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_26])])).
fof(c_0_55, plain, ![X42, X43, X44, X45]:(((r1_orders_2(X45,X44,X43)|~r1_lattice3(X45,k2_tarski(X43,X42),X44)|~epred1_4(X45,X44,X43,X42))&(r1_orders_2(X45,X44,X42)|~r1_lattice3(X45,k2_tarski(X43,X42),X44)|~epred1_4(X45,X44,X43,X42)))&((~r1_orders_2(X45,X44,X43)|~r1_orders_2(X45,X44,X42)|r1_lattice3(X45,k2_tarski(X43,X42),X44)|~epred1_4(X45,X44,X43,X42))&(((r1_orders_2(X45,X43,X44)|~r2_lattice3(X45,k2_tarski(X43,X42),X44)|~epred1_4(X45,X44,X43,X42))&(r1_orders_2(X45,X42,X44)|~r2_lattice3(X45,k2_tarski(X43,X42),X44)|~epred1_4(X45,X44,X43,X42)))&(~r1_orders_2(X45,X43,X44)|~r1_orders_2(X45,X42,X44)|r2_lattice3(X45,k2_tarski(X43,X42),X44)|~epred1_4(X45,X44,X43,X42))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_47])])])).
fof(c_0_56, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>epred1_4(X1,X2,X3,X4))))), inference(apply_def,[status(thm)],[t8_yellow_0, c_0_40])).
cnf(c_0_57, plain, (r1_orders_2(X1,X2,X3)|~r2_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_58, plain, (r1_yellow_0(X1,k2_tarski(X2,X3))|~epred3_3(X1,X2,X3)|~r1_orders_2(X1,X3,X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_59, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),k1_yellow_0(X1,X2))|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_24, c_0_31])).
cnf(c_0_60, plain, (k2_tarski(X2,X2)=k6_domain_1(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_61, negated_conjecture, (k1_yellow_0(esk1_0,X1)=esk2_0|~r1_yellow_0(esk1_0,X1)|~r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),esk2_0)|~r1_orders_2(esk1_0,esk2_0,k1_yellow_0(esk1_0,X1))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_25])])).
cnf(c_0_62, plain, (r1_orders_2(X1,X2,X3)|~r2_lattice3(X1,k2_tarski(X4,X2),X3)|~epred1_4(X1,X3,X4,X2)), inference(split_conjunct,[status(thm)],[c_0_55])).
fof(c_0_63, plain, ![X13, X14, X15, X16]:(~l1_orders_2(X13)|(~m1_subset_1(X14,u1_struct_0(X13))|(~m1_subset_1(X15,u1_struct_0(X13))|(~m1_subset_1(X16,u1_struct_0(X13))|epred1_4(X13,X14,X15,X16))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_56])])])).
cnf(c_0_64, plain, (r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r2_lattice3(X1,k2_tarski(X2,X2),X3)), inference(rw,[status(thm)],[c_0_57, c_0_52])).
cnf(c_0_65, plain, (r2_lattice3(X1,X2,esk4_3(X1,X2,X3))|X3=k1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_66, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X1))|X3=k1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_67, plain, (r1_yellow_0(X1,k2_tarski(X2,k1_yellow_0(X1,X3)))|v2_struct_0(X1)|~epred3_3(X1,X2,k1_yellow_0(X1,X3))|~r1_orders_2(X1,X2,k1_yellow_0(X1,X3))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_31])).
cnf(c_0_68, plain, (k2_tarski(k1_yellow_0(X1,X2),k1_yellow_0(X1,X2))=k6_domain_1(u1_struct_0(X1),k1_yellow_0(X1,X2))|v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_60, c_0_31])).
cnf(c_0_69, negated_conjecture, (k1_yellow_0(esk1_0,X1)=esk2_0|~r1_yellow_0(esk1_0,X1)|~r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),esk2_0)|~r1_orders_2(esk1_0,esk2_0,k1_yellow_0(esk1_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_31]), c_0_26])])).
cnf(c_0_70, plain, (r1_orders_2(X1,X2,k1_yellow_0(X1,k2_tarski(X3,X2)))|~epred1_4(X1,k1_yellow_0(X1,k2_tarski(X3,X2)),X3,X2)|~r1_yellow_0(X1,k2_tarski(X3,X2))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_62, c_0_38])).
cnf(c_0_71, plain, (epred1_4(X1,X2,X3,X4)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_72, plain, (X2=k1_yellow_0(X1,X3)|~r1_orders_2(X1,X2,esk4_3(X1,X3,X2))|~r2_lattice3(X1,X3,X2)|~r1_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_73, plain, (X1=k1_yellow_0(X2,k2_tarski(X3,X3))|r1_orders_2(X2,X3,esk4_3(X2,k2_tarski(X3,X3),X1))|~r1_yellow_0(X2,k2_tarski(X3,X3))|~r2_lattice3(X2,k2_tarski(X3,X3),X1)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_66])).
cnf(c_0_74, plain, (r2_lattice3(X1,k1_tarski(X2),X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_75, plain, ![X21, X22]:(~l1_orders_2(X21)|m1_subset_1(k2_yellow_0(X21,X22),u1_struct_0(X21))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_yellow_0])])).
cnf(c_0_76, plain, (v1_xboole_0(u1_struct_0(X1))|r1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),k1_yellow_0(X1,X2)))|v2_struct_0(X1)|~epred3_3(X1,k1_yellow_0(X1,X2),k1_yellow_0(X1,X2))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_59])).
cnf(c_0_77, negated_conjecture, (k1_yellow_0(esk1_0,X1)=esk2_0|~r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,k1_yellow_0(esk1_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_37]), c_0_25]), c_0_26])])).
cnf(c_0_78, plain, (r1_orders_2(X1,X2,k1_yellow_0(X1,k2_tarski(X3,X2)))|~r1_yellow_0(X1,k2_tarski(X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_31])).
cnf(c_0_79, plain, (k1_yellow_0(X1,k2_tarski(X2,X2))=X2|~r1_yellow_0(X1,k2_tarski(X2,X2))|~r2_lattice3(X1,k2_tarski(X2,X2),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_80, plain, (r2_lattice3(X1,k2_tarski(X2,X2),X3)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[c_0_74, c_0_52])).
cnf(c_0_81, plain, (r1_yellow_0(X1,k2_tarski(X2,X3))|r1_orders_2(X1,X2,esk6_4(X3,X2,X1,X4))|~r1_orders_2(X1,X3,X4)|~r1_orders_2(X1,X2,X4)|~m1_subset_1(X4,u1_struct_0(X1))|~epred3_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_82, plain, (m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_83, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),k1_yellow_0(esk1_0,X1)))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_54]), c_0_26]), c_0_27])]), c_0_28])).
cnf(c_0_84, plain, (k1_yellow_0(esk1_0,k2_tarski(X1,esk2_0))=esk2_0|~r1_yellow_0(esk1_0,k2_tarski(X1,esk2_0))|~r2_lattice3(esk1_0,k2_tarski(X1,esk2_0),esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_25]), c_0_26])])).
cnf(c_0_85, plain, (k1_yellow_0(X1,k2_tarski(X2,X2))=X2|~r1_yellow_0(X1,k2_tarski(X2,X2))|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_86, plain, (r1_yellow_0(X1,k2_tarski(X2,X3))|~epred3_3(X1,X2,X3)|~r1_orders_2(X1,X3,X2)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_49, c_0_81])).
cnf(c_0_87, plain, (r1_orders_2(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X2))|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_24, c_0_82])).
cnf(c_0_88, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))|~r1_yellow_0(esk1_0,k2_tarski(X1,esk2_0))|~r2_lattice3(esk1_0,k2_tarski(X1,esk2_0),esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_25])])).
cnf(c_0_89, plain, (r2_lattice3(X1,k2_tarski(X2,X4),X3)|~r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,X4,X3)|~epred1_4(X1,X3,X2,X4)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_90, negated_conjecture, (X1=esk2_0|~r1_yellow_0(esk1_0,k2_tarski(X1,X1))|~r1_orders_2(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_85]), c_0_26])])).
cnf(c_0_91, plain, (r1_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X2),X3))|v2_struct_0(X1)|~epred3_3(X1,k2_yellow_0(X1,X2),X3)|~r1_orders_2(X1,X3,k2_yellow_0(X1,X2))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_82])).
cnf(c_0_92, plain, (v1_xboole_0(u1_struct_0(esk1_0))|r1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))|~epred1_4(esk1_0,esk2_0,X1,esk2_0)|~r1_yellow_0(esk1_0,k2_tarski(X1,esk2_0))|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_35])])).
cnf(c_0_93, plain, (k2_yellow_0(esk1_0,X1)=esk2_0|~epred3_3(esk1_0,k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~r1_orders_2(esk1_0,esk2_0,k2_yellow_0(esk1_0,X1))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_26]), c_0_27])]), c_0_28])).
fof(c_0_94, plain, ![X8, X9, X10, X11]:(((r1_lattice3(X8,X9,X10)|X10!=k2_yellow_0(X8,X9)|~r2_yellow_0(X8,X9)|~m1_subset_1(X10,u1_struct_0(X8))|~l1_orders_2(X8))&(~m1_subset_1(X11,u1_struct_0(X8))|(~r1_lattice3(X8,X9,X11)|r1_orders_2(X8,X11,X10))|X10!=k2_yellow_0(X8,X9)|~r2_yellow_0(X8,X9)|~m1_subset_1(X10,u1_struct_0(X8))|~l1_orders_2(X8)))&((m1_subset_1(esk3_3(X8,X9,X10),u1_struct_0(X8))|~r1_lattice3(X8,X9,X10)|X10=k2_yellow_0(X8,X9)|~r2_yellow_0(X8,X9)|~m1_subset_1(X10,u1_struct_0(X8))|~l1_orders_2(X8))&((r1_lattice3(X8,X9,esk3_3(X8,X9,X10))|~r1_lattice3(X8,X9,X10)|X10=k2_yellow_0(X8,X9)|~r2_yellow_0(X8,X9)|~m1_subset_1(X10,u1_struct_0(X8))|~l1_orders_2(X8))&(~r1_orders_2(X8,esk3_3(X8,X9,X10),X10)|~r1_lattice3(X8,X9,X10)|X10=k2_yellow_0(X8,X9)|~r2_yellow_0(X8,X9)|~m1_subset_1(X10,u1_struct_0(X8))|~l1_orders_2(X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_yellow_0])])])])])).
cnf(c_0_95, plain, (v1_xboole_0(u1_struct_0(esk1_0))|r1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))|~r1_yellow_0(esk1_0,k2_tarski(X1,esk2_0))|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_71]), c_0_25]), c_0_26])])).
cnf(c_0_96, negated_conjecture, (r1_yellow_0(esk1_0,k2_tarski(X1,esk2_0))|~epred3_3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_35]), c_0_25])])).
cnf(c_0_97, negated_conjecture, (k2_yellow_0(esk1_0,X1)=esk2_0|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~r1_orders_2(esk1_0,esk2_0,k2_yellow_0(esk1_0,X1))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_93, c_0_54])).
cnf(c_0_98, plain, (r1_orders_2(X1,X3,X2)|~r1_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_99, plain, (r1_lattice3(X1,X2,X3)|X3!=k2_yellow_0(X1,X2)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_100, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))|~epred3_3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_95, c_0_96])).
fof(c_0_101, plain, ![X31]:(v2_struct_0(X31)|~l1_struct_0(X31)|~v1_xboole_0(u1_struct_0(X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_102, negated_conjecture, (k2_yellow_0(esk1_0,X1)=esk2_0|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~r1_orders_2(esk1_0,esk2_0,k2_yellow_0(esk1_0,X1))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_87]), c_0_26]), c_0_27])]), c_0_28])).
cnf(c_0_103, plain, (r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,k2_tarski(X2,X2),X3)), inference(rw,[status(thm)],[c_0_98, c_0_52])).
cnf(c_0_104, plain, (r1_lattice3(X1,X2,k2_yellow_0(X1,X2))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_99]), c_0_82])).
cnf(c_0_105, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_54]), c_0_25])])).
cnf(c_0_106, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_107, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_60, c_0_25])).
cnf(c_0_108, negated_conjecture, (k2_yellow_0(esk1_0,X1)=esk2_0|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~r1_orders_2(esk1_0,esk2_0,k2_yellow_0(esk1_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_82]), c_0_26])])).
cnf(c_0_109, plain, (r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(X2,X2)),X2)|~r2_yellow_0(X1,k2_tarski(X2,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_82])).
cnf(c_0_110, plain, (r1_orders_2(X2,X1,X4)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|X4!=k2_yellow_0(X2,X3)|~r2_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_111, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_25]), c_0_35])])).
cnf(c_0_112, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_28])).
cnf(c_0_113, negated_conjecture, (k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_114, negated_conjecture, (k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))=esk2_0|~r1_orders_2(esk1_0,esk2_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_25]), c_0_26])])).
cnf(c_0_115, plain, (r1_orders_2(X1,X2,k2_yellow_0(X1,X3))|~r1_lattice3(X1,X3,X2)|~r2_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_110]), c_0_82])).
cnf(c_0_116, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_111, c_0_112])).
fof(c_0_117, plain, ![X17]:(~l1_orders_2(X17)|l1_struct_0(X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
fof(c_0_118, plain, ![X3, X2, X1]:(epred2_3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))), introduced(definition)).
cnf(c_0_119, negated_conjecture, (k1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))!=esk2_0|~l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_113, c_0_112])).
cnf(c_0_120, negated_conjecture, (k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))=esk2_0|~r1_lattice3(esk1_0,k2_tarski(esk2_0,esk2_0),esk2_0)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_25]), c_0_26])])).
cnf(c_0_121, negated_conjecture, (r1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_116]), c_0_28])).
cnf(c_0_122, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_117])).
fof(c_0_123, plain, ![X3, X2, X1]:(epred2_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))), inference(split_equiv,[status(thm)],[c_0_118])).
cnf(c_0_124, negated_conjecture, (k1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))!=esk2_0|~l1_struct_0(esk1_0)|~r1_lattice3(esk1_0,k2_tarski(esk2_0,esk2_0),esk2_0)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(spm,[status(thm)],[c_0_119, c_0_120])).
cnf(c_0_125, negated_conjecture, (r1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_26])])).
cnf(c_0_126, plain, (r1_lattice3(X1,k1_tarski(X3),X2)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_127, plain, ![X46, X47, X48, X49, X50]:(((r1_orders_2(X48,X49,X47)|(X49!=k11_lattice3(X48,X47,X46)|~r2_yellow_0(X48,k2_tarski(X47,X46)))|~m1_subset_1(X49,u1_struct_0(X48))|~epred2_3(X48,X47,X46))&((r1_orders_2(X48,X49,X46)|(X49!=k11_lattice3(X48,X47,X46)|~r2_yellow_0(X48,k2_tarski(X47,X46)))|~m1_subset_1(X49,u1_struct_0(X48))|~epred2_3(X48,X47,X46))&(~m1_subset_1(X50,u1_struct_0(X48))|(~r1_orders_2(X48,X50,X47)|~r1_orders_2(X48,X50,X46)|r1_orders_2(X48,X50,X49))|(X49!=k11_lattice3(X48,X47,X46)|~r2_yellow_0(X48,k2_tarski(X47,X46)))|~m1_subset_1(X49,u1_struct_0(X48))|~epred2_3(X48,X47,X46))))&(((X49=k11_lattice3(X48,X47,X46)|(m1_subset_1(esk5_4(X46,X47,X48,X49),u1_struct_0(X48))|~r1_orders_2(X48,X49,X46)|~r1_orders_2(X48,X49,X47))|~m1_subset_1(X49,u1_struct_0(X48))|~epred2_3(X48,X47,X46))&(r2_yellow_0(X48,k2_tarski(X47,X46))|(m1_subset_1(esk5_4(X46,X47,X48,X49),u1_struct_0(X48))|~r1_orders_2(X48,X49,X46)|~r1_orders_2(X48,X49,X47))|~m1_subset_1(X49,u1_struct_0(X48))|~epred2_3(X48,X47,X46)))&((((X49=k11_lattice3(X48,X47,X46)|(r1_orders_2(X48,esk5_4(X46,X47,X48,X49),X47)|~r1_orders_2(X48,X49,X46)|~r1_orders_2(X48,X49,X47))|~m1_subset_1(X49,u1_struct_0(X48))|~epred2_3(X48,X47,X46))&(r2_yellow_0(X48,k2_tarski(X47,X46))|(r1_orders_2(X48,esk5_4(X46,X47,X48,X49),X47)|~r1_orders_2(X48,X49,X46)|~r1_orders_2(X48,X49,X47))|~m1_subset_1(X49,u1_struct_0(X48))|~epred2_3(X48,X47,X46)))&((X49=k11_lattice3(X48,X47,X46)|(r1_orders_2(X48,esk5_4(X46,X47,X48,X49),X46)|~r1_orders_2(X48,X49,X46)|~r1_orders_2(X48,X49,X47))|~m1_subset_1(X49,u1_struct_0(X48))|~epred2_3(X48,X47,X46))&(r2_yellow_0(X48,k2_tarski(X47,X46))|(r1_orders_2(X48,esk5_4(X46,X47,X48,X49),X46)|~r1_orders_2(X48,X49,X46)|~r1_orders_2(X48,X49,X47))|~m1_subset_1(X49,u1_struct_0(X48))|~epred2_3(X48,X47,X46))))&((X49=k11_lattice3(X48,X47,X46)|(~r1_orders_2(X48,esk5_4(X46,X47,X48,X49),X49)|~r1_orders_2(X48,X49,X46)|~r1_orders_2(X48,X49,X47))|~m1_subset_1(X49,u1_struct_0(X48))|~epred2_3(X48,X47,X46))&(r2_yellow_0(X48,k2_tarski(X47,X46))|(~r1_orders_2(X48,esk5_4(X46,X47,X48,X49),X49)|~r1_orders_2(X48,X49,X46)|~r1_orders_2(X48,X49,X47))|~m1_subset_1(X49,u1_struct_0(X48))|~epred2_3(X48,X47,X46)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_123])])])])])).
cnf(c_0_128, negated_conjecture, (~l1_struct_0(esk1_0)|~r1_lattice3(esk1_0,k2_tarski(esk2_0,esk2_0),esk2_0)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_85]), c_0_125]), c_0_35]), c_0_25]), c_0_26])])).
cnf(c_0_129, plain, (r1_lattice3(X1,k2_tarski(X3,X3),X2)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[c_0_126, c_0_52])).
cnf(c_0_130, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|~r1_orders_2(X1,esk5_4(X3,X2,X1,X4),X4)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~epred2_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_131, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|r1_orders_2(X1,esk5_4(X3,X2,X1,X4),X3)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~epred2_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_127])).
fof(c_0_132, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>epred2_3(X1,X2,X3)))), inference(apply_def,[status(thm)],[t19_yellow_0, c_0_118])).
cnf(c_0_133, negated_conjecture, (~l1_struct_0(esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129]), c_0_35]), c_0_25]), c_0_26])])).
cnf(c_0_134, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|~epred2_3(X1,X2,X3)|~r1_orders_2(X1,X3,X3)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_130, c_0_131])).
fof(c_0_135, plain, ![X28, X29, X30]:(~v5_orders_2(X28)|~l1_orders_2(X28)|(~m1_subset_1(X29,u1_struct_0(X28))|(~m1_subset_1(X30,u1_struct_0(X28))|epred2_3(X28,X29,X30)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_132])])])).
cnf(c_0_136, negated_conjecture, (~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_122]), c_0_26])])).
cnf(c_0_137, negated_conjecture, (r2_yellow_0(esk1_0,k2_tarski(X1,esk2_0))|~epred2_3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_35]), c_0_25])])).
cnf(c_0_138, plain, (epred2_3(X1,X2,X3)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_135])).
cnf(c_0_139, negated_conjecture, (~epred2_3(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_137]), c_0_35])])).
cnf(c_0_140, negated_conjecture, (epred2_3(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_46]), c_0_26])])).
cnf(c_0_141, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140]), c_0_25])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 142
# Proof object clause steps            : 101
# Proof object formula steps           : 41
# Proof object conjectures             : 42
# Proof object clause conjectures      : 39
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 37
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 55
# Proof object simplifying inferences  : 99
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 58
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 57
# Processed clauses                    : 411
# ...of these trivial                  : 0
# ...subsumed                          : 115
# ...remaining for further processing  : 296
# Other redundant clauses eliminated   : 10
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 29
# Backward-rewritten                   : 4
# Generated clauses                    : 790
# ...of the previous two non-trivial   : 659
# Contextual simplify-reflections      : 47
# Paramodulations                      : 780
# Factorizations                       : 0
# Equation resolutions                 : 10
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 196
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 186
# Current number of unprocessed clauses: 350
# ...number of literals in the above   : 2933
# Current number of archived formulas  : 0
# Current number of archived clauses   : 91
# Clause-clause subsumption calls (NU) : 8302
# Rec. Clause-clause subsumption calls : 652
# Non-unit clause-clause subsumptions  : 190
# Unit Clause-clause subsumption calls : 146
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 29819

# -------------------------------------------------
# User time                : 0.032 s
# System time              : 0.000 s
# Total time               : 0.032 s
# Maximum resident set size: 28720 pages
