# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t104_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(u1_struct_0(k6_tmap_1(X1,X2))=u1_struct_0(X1)&u1_pre_topc(k6_tmap_1(X1,X2))=k5_tmap_1(X1,X2)))), file('data/MPTP/proofs/22077886fd8c485c9b7e8a1bddded455.E_input', t104_tmap_1)).
fof(t1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/22077886fd8c485c9b7e8a1bddded455.E_input', t1_tsep_1)).
fof(d11_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:((v1_pre_topc(X3)&(v2_pre_topc(X3)&l1_pre_topc(X3)))=>(X3=k8_tmap_1(X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(X4=u1_struct_0(X2)=>X3=k6_tmap_1(X1,X4))))))), file('data/MPTP/proofs/22077886fd8c485c9b7e8a1bddded455.E_input', d11_tmap_1)).
fof(dt_k8_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_pre_topc(X2,X1))=>(v1_pre_topc(k8_tmap_1(X1,X2))&(v2_pre_topc(k8_tmap_1(X1,X2))&l1_pre_topc(k8_tmap_1(X1,X2))))), file('data/MPTP/proofs/22077886fd8c485c9b7e8a1bddded455.E_input', dt_k8_tmap_1)).
fof(t119_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>r1_tmap_1(X2,k8_tmap_1(X1,X2),k2_tmap_1(X1,k8_tmap_1(X1,X2),k9_tmap_1(X1,X2),X2),X3)))), file('data/MPTP/proofs/22077886fd8c485c9b7e8a1bddded455.E_input', t119_tmap_1)).
fof(t2_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>m1_pre_topc(X1,X1)), file('data/MPTP/proofs/22077886fd8c485c9b7e8a1bddded455.E_input', t2_tsep_1)).
fof(d10_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k7_tmap_1(X1,X2)=k6_partfun1(u1_struct_0(X1)))), file('data/MPTP/proofs/22077886fd8c485c9b7e8a1bddded455.E_input', d10_tmap_1)).
fof(dt_k7_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_funct_1(k7_tmap_1(X1,X2))&(v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))&m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))))), file('data/MPTP/proofs/22077886fd8c485c9b7e8a1bddded455.E_input', dt_k7_tmap_1)).
fof(d12_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))))=>(X3=k9_tmap_1(X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(X4=u1_struct_0(X2)=>r1_funct_2(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X4)),X3,k7_tmap_1(X1,X4)))))))), file('data/MPTP/proofs/22077886fd8c485c9b7e8a1bddded455.E_input', d12_tmap_1)).
fof(t110_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(u1_struct_0(X3)=X2=>![X4]:(m1_subset_1(X4,u1_struct_0(X3))=>r1_tmap_1(X3,k6_tmap_1(X1,X2),k2_tmap_1(X1,k6_tmap_1(X1,X2),k7_tmap_1(X1,X2),X3),X4)))))), file('data/MPTP/proofs/22077886fd8c485c9b7e8a1bddded455.E_input', t110_tmap_1)).
fof(reflexivity_r1_funct_2, axiom, ![X1, X2, X3, X4, X5, X6]:((~(v1_xboole_0(X2))&(~(v1_xboole_0(X4))&((v1_funct_1(X5)&(v1_funct_2(X5,X1,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&(v1_funct_1(X6)&(v1_funct_2(X6,X3,X4)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))))))=>r1_funct_2(X1,X2,X3,X4,X5,X5)), file('data/MPTP/proofs/22077886fd8c485c9b7e8a1bddded455.E_input', reflexivity_r1_funct_2)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/22077886fd8c485c9b7e8a1bddded455.E_input', fc2_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/22077886fd8c485c9b7e8a1bddded455.E_input', dt_l1_pre_topc)).
fof(c_0_13, plain, ![X11, X12]:((u1_struct_0(k6_tmap_1(X11,X12))=u1_struct_0(X11)|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|(v2_struct_0(X11)|(~v2_pre_topc(X11)|~l1_pre_topc(X11))))&(u1_pre_topc(k6_tmap_1(X11,X12))=k5_tmap_1(X11,X12)|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|(v2_struct_0(X11)|(~v2_pre_topc(X11)|~l1_pre_topc(X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t104_tmap_1])])])])])).
fof(c_0_14, plain, ![X29, X30]:(~l1_pre_topc(X29)|(~m1_pre_topc(X30,X29)|m1_subset_1(u1_struct_0(X30),k1_zfmisc_1(u1_struct_0(X29))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])).
fof(c_0_15, plain, ![X38, X39, X40, X41]:((X40!=k8_tmap_1(X38,X39)|(~m1_subset_1(X41,k1_zfmisc_1(u1_struct_0(X38)))|(X41!=u1_struct_0(X39)|X40=k6_tmap_1(X38,X41)))|(~v1_pre_topc(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40)))|~m1_pre_topc(X39,X38)|(v2_struct_0(X38)|(~v2_pre_topc(X38)|~l1_pre_topc(X38))))&((m1_subset_1(esk5_3(X38,X39,X40),k1_zfmisc_1(u1_struct_0(X38)))|X40=k8_tmap_1(X38,X39)|(~v1_pre_topc(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40)))|~m1_pre_topc(X39,X38)|(v2_struct_0(X38)|(~v2_pre_topc(X38)|~l1_pre_topc(X38))))&((esk5_3(X38,X39,X40)=u1_struct_0(X39)|X40=k8_tmap_1(X38,X39)|(~v1_pre_topc(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40)))|~m1_pre_topc(X39,X38)|(v2_struct_0(X38)|(~v2_pre_topc(X38)|~l1_pre_topc(X38))))&(X40!=k6_tmap_1(X38,esk5_3(X38,X39,X40))|X40=k8_tmap_1(X38,X39)|(~v1_pre_topc(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40)))|~m1_pre_topc(X39,X38)|(v2_struct_0(X38)|(~v2_pre_topc(X38)|~l1_pre_topc(X38))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d11_tmap_1])])])])])])).
fof(c_0_16, plain, ![X36, X37]:((v1_pre_topc(k8_tmap_1(X36,X37))|(v2_struct_0(X36)|(~v2_pre_topc(X36)|~l1_pre_topc(X36))|~m1_pre_topc(X37,X36)))&((v2_pre_topc(k8_tmap_1(X36,X37))|(v2_struct_0(X36)|(~v2_pre_topc(X36)|~l1_pre_topc(X36))|~m1_pre_topc(X37,X36)))&(l1_pre_topc(k8_tmap_1(X36,X37))|(v2_struct_0(X36)|(~v2_pre_topc(X36)|~l1_pre_topc(X36))|~m1_pre_topc(X37,X36))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k8_tmap_1])])])])).
cnf(c_0_17, plain, (u1_struct_0(k6_tmap_1(X1,X2))=u1_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (X1=k6_tmap_1(X2,X4)|v2_struct_0(X2)|X1!=k8_tmap_1(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|X4!=u1_struct_0(X3)|~v1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, plain, (v2_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, plain, (l1_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (v1_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_23, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>r1_tmap_1(X2,k8_tmap_1(X1,X2),k2_tmap_1(X1,k8_tmap_1(X1,X2),k9_tmap_1(X1,X2),X2),X3))))), inference(assume_negation,[status(cth)],[t119_tmap_1])).
cnf(c_0_24, plain, (u1_struct_0(k6_tmap_1(X1,u1_struct_0(X2)))=u1_struct_0(X1)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_25, plain, (k6_tmap_1(X1,u1_struct_0(X2))=k8_tmap_1(X1,X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_19])]), c_0_20]), c_0_21]), c_0_18]), c_0_22])).
fof(c_0_26, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk2_0))&~r1_tmap_1(esk2_0,k8_tmap_1(esk1_0,esk2_0),k2_tmap_1(esk1_0,k8_tmap_1(esk1_0,esk2_0),k9_tmap_1(esk1_0,esk2_0),esk2_0),esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_23])])])])).
cnf(c_0_27, plain, (u1_struct_0(k8_tmap_1(X1,X2))=u1_struct_0(X1)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_28, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_29, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_30, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_31, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, negated_conjecture, (u1_struct_0(k8_tmap_1(esk1_0,esk2_0))=u1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30])]), c_0_31])).
fof(c_0_33, plain, ![X19]:(~l1_pre_topc(X19)|m1_pre_topc(X19,X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tsep_1])])).
fof(c_0_34, plain, ![X27, X28]:(v2_struct_0(X27)|(~v2_pre_topc(X27)|~l1_pre_topc(X27))|(~m1_subset_1(X28,k1_zfmisc_1(u1_struct_0(X27)))|k7_tmap_1(X27,X28)=k6_partfun1(u1_struct_0(X27)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_tmap_1])])])])).
cnf(c_0_35, negated_conjecture, (m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_pre_topc(k8_tmap_1(esk1_0,esk2_0),X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_18, c_0_32])).
cnf(c_0_36, plain, (m1_pre_topc(X1,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_37, plain, (v2_struct_0(X1)|k7_tmap_1(X1,X2)=k6_partfun1(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_38, negated_conjecture, (m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_32])).
fof(c_0_39, plain, ![X20, X21]:((v1_funct_1(k7_tmap_1(X20,X21))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))))&((v1_funct_2(k7_tmap_1(X20,X21),u1_struct_0(X20),u1_struct_0(k6_tmap_1(X20,X21)))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))))&(m1_subset_1(k7_tmap_1(X20,X21),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X20),u1_struct_0(k6_tmap_1(X20,X21)))))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_tmap_1])])])])).
cnf(c_0_40, negated_conjecture, (k6_partfun1(u1_struct_0(esk1_0))=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|~l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_41, plain, (k7_tmap_1(X1,u1_struct_0(X2))=k6_partfun1(u1_struct_0(X1))|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_37, c_0_18])).
fof(c_0_42, plain, ![X31, X32, X33, X34]:((X33!=k9_tmap_1(X31,X32)|(~m1_subset_1(X34,k1_zfmisc_1(u1_struct_0(X31)))|(X34!=u1_struct_0(X32)|r1_funct_2(u1_struct_0(X31),u1_struct_0(k8_tmap_1(X31,X32)),u1_struct_0(X31),u1_struct_0(k6_tmap_1(X31,X34)),X33,k7_tmap_1(X31,X34))))|(~v1_funct_1(X33)|(~v1_funct_2(X33,u1_struct_0(X31),u1_struct_0(k8_tmap_1(X31,X32)))|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X31),u1_struct_0(k8_tmap_1(X31,X32)))))))|~m1_pre_topc(X32,X31)|(v2_struct_0(X31)|(~v2_pre_topc(X31)|~l1_pre_topc(X31))))&((m1_subset_1(esk4_3(X31,X32,X33),k1_zfmisc_1(u1_struct_0(X31)))|X33=k9_tmap_1(X31,X32)|(~v1_funct_1(X33)|(~v1_funct_2(X33,u1_struct_0(X31),u1_struct_0(k8_tmap_1(X31,X32)))|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X31),u1_struct_0(k8_tmap_1(X31,X32)))))))|~m1_pre_topc(X32,X31)|(v2_struct_0(X31)|(~v2_pre_topc(X31)|~l1_pre_topc(X31))))&((esk4_3(X31,X32,X33)=u1_struct_0(X32)|X33=k9_tmap_1(X31,X32)|(~v1_funct_1(X33)|(~v1_funct_2(X33,u1_struct_0(X31),u1_struct_0(k8_tmap_1(X31,X32)))|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X31),u1_struct_0(k8_tmap_1(X31,X32)))))))|~m1_pre_topc(X32,X31)|(v2_struct_0(X31)|(~v2_pre_topc(X31)|~l1_pre_topc(X31))))&(~r1_funct_2(u1_struct_0(X31),u1_struct_0(k8_tmap_1(X31,X32)),u1_struct_0(X31),u1_struct_0(k6_tmap_1(X31,esk4_3(X31,X32,X33))),X33,k7_tmap_1(X31,esk4_3(X31,X32,X33)))|X33=k9_tmap_1(X31,X32)|(~v1_funct_1(X33)|(~v1_funct_2(X33,u1_struct_0(X31),u1_struct_0(k8_tmap_1(X31,X32)))|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X31),u1_struct_0(k8_tmap_1(X31,X32)))))))|~m1_pre_topc(X32,X31)|(v2_struct_0(X31)|(~v2_pre_topc(X31)|~l1_pre_topc(X31))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d12_tmap_1])])])])])])).
cnf(c_0_43, plain, (m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_44, plain, (v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_45, plain, (v1_funct_1(k7_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_46, negated_conjecture, (k7_tmap_1(esk1_0,u1_struct_0(X1))=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|~m1_pre_topc(X1,esk1_0)|~l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_47, plain, (esk4_3(X1,X2,X3)=u1_struct_0(X2)|X3=k9_tmap_1(X1,X2)|v2_struct_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_48, plain, (m1_subset_1(k7_tmap_1(X1,u1_struct_0(X2)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_25]), c_0_18])).
cnf(c_0_49, plain, (v1_funct_2(k7_tmap_1(X1,u1_struct_0(X2)),u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_25]), c_0_18])).
cnf(c_0_50, negated_conjecture, (v1_funct_1(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X1,esk1_0)|~l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_51, plain, (m1_subset_1(esk4_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|X3=k9_tmap_1(X1,X2)|v2_struct_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_52, negated_conjecture, (esk4_3(esk1_0,esk2_0,X1)=u1_struct_0(esk2_0)|X1=k9_tmap_1(esk1_0,esk2_0)|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_32]), c_0_28]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_53, negated_conjecture, (m1_subset_1(k7_tmap_1(esk1_0,u1_struct_0(esk2_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_32]), c_0_28]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_54, negated_conjecture, (v1_funct_2(k7_tmap_1(esk1_0,u1_struct_0(esk2_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_32]), c_0_28]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_55, negated_conjecture, (v1_funct_1(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))|~m1_pre_topc(X1,esk1_0)|~l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_18]), c_0_29])])).
cnf(c_0_56, negated_conjecture, (X1=k9_tmap_1(esk1_0,esk2_0)|m1_subset_1(esk4_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_32]), c_0_28]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_57, negated_conjecture, (esk4_3(esk1_0,esk2_0,k7_tmap_1(esk1_0,u1_struct_0(esk2_0)))=u1_struct_0(esk2_0)|k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(esk2_0))|~v1_funct_1(k7_tmap_1(esk1_0,u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54])])).
cnf(c_0_58, negated_conjecture, (v1_funct_1(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))|~l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_55, c_0_28])).
cnf(c_0_59, negated_conjecture, (k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(esk2_0))|m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_1(k7_tmap_1(esk1_0,u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_54]), c_0_53])])).
cnf(c_0_60, negated_conjecture, (v1_funct_1(k7_tmap_1(esk1_0,u1_struct_0(X1)))|~m1_pre_topc(X1,esk1_0)|~l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_58, c_0_46])).
cnf(c_0_61, negated_conjecture, (k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(esk2_0))|m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_28])])).
cnf(c_0_62, negated_conjecture, (k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_46]), c_0_28])]), c_0_58])).
cnf(c_0_63, negated_conjecture, (k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(esk2_0))|k6_partfun1(u1_struct_0(esk1_0))=k7_tmap_1(esk1_0,u1_struct_0(esk2_0))|~l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_61]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_64, negated_conjecture, (k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|k6_partfun1(u1_struct_0(esk1_0))=k7_tmap_1(esk1_0,u1_struct_0(esk2_0))|~l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_62]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_65, negated_conjecture, (k7_tmap_1(esk1_0,u1_struct_0(X1))=k7_tmap_1(esk1_0,u1_struct_0(esk2_0))|k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(esk2_0))|~m1_pre_topc(X1,esk1_0)|~l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_63]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_66, negated_conjecture, (k7_tmap_1(esk1_0,u1_struct_0(X1))=k7_tmap_1(esk1_0,u1_struct_0(esk2_0))|k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|~m1_pre_topc(X1,esk1_0)|~l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_64]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_67, negated_conjecture, (k7_tmap_1(esk1_0,u1_struct_0(X1))=k7_tmap_1(esk1_0,u1_struct_0(esk2_0))|k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(esk2_0))|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_21]), c_0_28]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_68, negated_conjecture, (k7_tmap_1(esk1_0,u1_struct_0(X1))=k7_tmap_1(esk1_0,u1_struct_0(esk2_0))|k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_21]), c_0_28]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_69, negated_conjecture, (k7_tmap_1(esk1_0,u1_struct_0(esk2_0))=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_36]), c_0_29])])).
cnf(c_0_70, negated_conjecture, (k7_tmap_1(esk1_0,u1_struct_0(esk2_0))=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_36]), c_0_29])])).
cnf(c_0_71, negated_conjecture, (k7_tmap_1(esk1_0,u1_struct_0(esk2_0))=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_72, plain, (k7_tmap_1(X1,u1_struct_0(X2))=k7_tmap_1(X1,u1_struct_0(X3))|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_41, c_0_41])).
cnf(c_0_73, plain, (u1_struct_0(k8_tmap_1(X1,X1))=u1_struct_0(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_27, c_0_36])).
cnf(c_0_74, negated_conjecture, (k7_tmap_1(esk1_0,u1_struct_0(X1))=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_28]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_75, plain, (esk4_3(X1,X1,X2)=u1_struct_0(X1)|X2=k9_tmap_1(X1,X1)|v2_struct_0(X1)|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_73]), c_0_36])).
cnf(c_0_76, negated_conjecture, (v1_funct_1(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_74]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_77, plain, (X1=k9_tmap_1(X2,X2)|m1_subset_1(esk4_3(X2,X2,X1),k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2))))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_73]), c_0_36])).
cnf(c_0_78, negated_conjecture, (esk4_3(esk1_0,esk1_0,k7_tmap_1(esk1_0,u1_struct_0(esk2_0)))=u1_struct_0(esk1_0)|k9_tmap_1(esk1_0,esk1_0)=k7_tmap_1(esk1_0,u1_struct_0(esk2_0))|~v1_funct_1(k7_tmap_1(esk1_0,u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_53]), c_0_54]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_79, negated_conjecture, (k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_1(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_71]), c_0_71])).
cnf(c_0_80, negated_conjecture, (v1_funct_1(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))|~m1_pre_topc(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_18]), c_0_29])])).
cnf(c_0_81, negated_conjecture, (k9_tmap_1(esk1_0,esk1_0)=k7_tmap_1(esk1_0,u1_struct_0(esk2_0))|m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_1(k7_tmap_1(esk1_0,u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_54]), c_0_53]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_82, negated_conjecture, (k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(X1))|m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X1,esk1_0)|~l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_46]), c_0_60])).
fof(c_0_83, plain, ![X22, X23, X24, X25]:(v2_struct_0(X22)|(~v2_pre_topc(X22)|~l1_pre_topc(X22))|(~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(X22)))|(v2_struct_0(X24)|~m1_pre_topc(X24,X22)|(u1_struct_0(X24)!=X23|(~m1_subset_1(X25,u1_struct_0(X24))|r1_tmap_1(X24,k6_tmap_1(X22,X23),k2_tmap_1(X22,k6_tmap_1(X22,X23),k7_tmap_1(X22,X23),X24),X25)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t110_tmap_1])])])])).
cnf(c_0_84, plain, (X3=k9_tmap_1(X1,X2)|v2_struct_0(X1)|~r1_funct_2(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,esk4_3(X1,X2,X3))),X3,k7_tmap_1(X1,esk4_3(X1,X2,X3)))|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_85, negated_conjecture, (esk4_3(esk1_0,esk1_0,k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))=u1_struct_0(esk1_0)|k9_tmap_1(esk1_0,esk1_0)=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|~v1_funct_1(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_71]), c_0_71]), c_0_71])).
cnf(c_0_86, negated_conjecture, (v1_funct_1(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_80, c_0_28])).
fof(c_0_87, plain, ![X13, X14, X15, X16, X17, X18]:(v1_xboole_0(X14)|(v1_xboole_0(X16)|(~v1_funct_1(X17)|(~v1_funct_2(X17,X13,X14)|~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))|(~v1_funct_1(X18)|(~v1_funct_2(X18,X15,X16)|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(X15,X16)))))))|r1_funct_2(X13,X14,X15,X16,X17,X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r1_funct_2])])])).
cnf(c_0_88, negated_conjecture, (k9_tmap_1(esk1_0,esk1_0)=k7_tmap_1(esk1_0,u1_struct_0(esk2_0))|m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_45]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_89, negated_conjecture, (u1_struct_0(k6_tmap_1(esk1_0,u1_struct_0(esk1_0)))=u1_struct_0(esk1_0)|~l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_38]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_90, negated_conjecture, (k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(X1))|m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_21]), c_0_28]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_91, plain, (v2_struct_0(X1)|v2_struct_0(X3)|r1_tmap_1(X3,k6_tmap_1(X1,X2),k2_tmap_1(X1,k6_tmap_1(X1,X2),k7_tmap_1(X1,X2),X3),X4)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_pre_topc(X3,X1)|u1_struct_0(X3)!=X2|~m1_subset_1(X4,u1_struct_0(X3))), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_92, plain, (X1=k9_tmap_1(X2,X2)|v2_struct_0(X2)|~r1_funct_2(u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(k6_tmap_1(X2,esk4_3(X2,X2,X1))),X1,k7_tmap_1(X2,esk4_3(X2,X2,X1)))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2))))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_73]), c_0_36])).
cnf(c_0_93, negated_conjecture, (esk4_3(esk1_0,esk1_0,k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))=u1_struct_0(esk1_0)|k9_tmap_1(esk1_0,esk1_0)=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_86])])).
cnf(c_0_94, negated_conjecture, (v1_funct_2(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(rw,[status(thm)],[c_0_54, c_0_71])).
cnf(c_0_95, negated_conjecture, (m1_subset_1(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(rw,[status(thm)],[c_0_53, c_0_71])).
cnf(c_0_96, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|r1_funct_2(X4,X1,X6,X2,X3,X3)|~v1_funct_1(X3)|~v1_funct_2(X3,X4,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X1)))|~v1_funct_1(X5)|~v1_funct_2(X5,X6,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X2)))), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_97, negated_conjecture, (k9_tmap_1(esk1_0,esk1_0)=k7_tmap_1(esk1_0,u1_struct_0(esk2_0))|m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_18]), c_0_28]), c_0_29])])).
cnf(c_0_98, negated_conjecture, (u1_struct_0(k8_tmap_1(esk1_0,esk1_0))=u1_struct_0(esk1_0)|~m1_pre_topc(esk1_0,esk1_0)|~l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_25]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_99, negated_conjecture, (k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_28]), c_0_71])).
cnf(c_0_100, plain, (r1_tmap_1(X1,k6_tmap_1(X2,u1_struct_0(X1)),k2_tmap_1(X2,k6_tmap_1(X2,u1_struct_0(X1)),k7_tmap_1(X2,u1_struct_0(X1)),X1),X3)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_91]), c_0_18])).
fof(c_0_101, plain, ![X10]:(v2_struct_0(X10)|~l1_struct_0(X10)|~v1_xboole_0(u1_struct_0(X10))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_102, plain, ![X26]:(~l1_pre_topc(X26)|l1_struct_0(X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_103, negated_conjecture, (k9_tmap_1(esk1_0,esk1_0)=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|~r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,u1_struct_0(esk1_0))),k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_94]), c_0_86]), c_0_95]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_104, plain, (r1_funct_2(X1,X2,u1_struct_0(X3),u1_struct_0(k6_tmap_1(X3,X4)),X5,X5)|v1_xboole_0(u1_struct_0(k6_tmap_1(X3,X4)))|v1_xboole_0(X2)|v2_struct_0(X3)|~v1_funct_2(X5,X1,X2)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3)|~v2_pre_topc(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_43]), c_0_45]), c_0_44])).
cnf(c_0_105, negated_conjecture, (k9_tmap_1(esk1_0,esk1_0)=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(rw,[status(thm)],[c_0_97, c_0_71])).
cnf(c_0_106, negated_conjecture, (u1_struct_0(k8_tmap_1(esk1_0,esk1_0))=u1_struct_0(esk1_0)|~m1_pre_topc(esk1_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_21]), c_0_28]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_107, negated_conjecture, (~r1_tmap_1(esk2_0,k8_tmap_1(esk1_0,esk2_0),k2_tmap_1(esk1_0,k8_tmap_1(esk1_0,esk2_0),k9_tmap_1(esk1_0,esk2_0),esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_108, negated_conjecture, (k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|u1_struct_0(k6_tmap_1(esk1_0,u1_struct_0(esk2_0)))=u1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_99]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_109, plain, (r1_tmap_1(X1,k8_tmap_1(X2,X1),k2_tmap_1(X2,k8_tmap_1(X2,X1),k7_tmap_1(X2,u1_struct_0(X1)),X1),X3)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_100, c_0_25])).
cnf(c_0_110, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_111, plain, (r1_funct_2(u1_struct_0(X2),u1_struct_0(k8_tmap_1(X2,X3)),u1_struct_0(X2),u1_struct_0(k6_tmap_1(X2,X4)),X1,k7_tmap_1(X2,X4))|v2_struct_0(X2)|X1!=k9_tmap_1(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|X4!=u1_struct_0(X3)|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(k8_tmap_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k8_tmap_1(X2,X3)))))|~m1_pre_topc(X3,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_112, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_113, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_114, negated_conjecture, (k9_tmap_1(esk1_0,esk1_0)=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(k6_tmap_1(esk1_0,u1_struct_0(esk1_0))))|v1_xboole_0(u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_94]), c_0_86]), c_0_95]), c_0_29]), c_0_30])]), c_0_31]), c_0_105])).
cnf(c_0_115, negated_conjecture, (u1_struct_0(k8_tmap_1(esk1_0,esk1_0))=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_36]), c_0_29])])).
cnf(c_0_116, negated_conjecture, (esk4_3(esk1_0,esk2_0,k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))=u1_struct_0(esk2_0)|k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|~v1_funct_1(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_71]), c_0_71]), c_0_71])).
cnf(c_0_117, negated_conjecture, (u1_struct_0(k6_tmap_1(esk1_0,u1_struct_0(esk2_0)))=u1_struct_0(esk1_0)|~r1_tmap_1(esk2_0,k8_tmap_1(esk1_0,esk2_0),k2_tmap_1(esk1_0,k8_tmap_1(esk1_0,esk2_0),k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_118, negated_conjecture, (r1_tmap_1(esk2_0,k8_tmap_1(esk1_0,esk2_0),k2_tmap_1(esk1_0,k8_tmap_1(esk1_0,esk2_0),k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),esk2_0),X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_71]), c_0_28]), c_0_29]), c_0_30])]), c_0_110]), c_0_31])).
cnf(c_0_119, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_120, plain, (r1_funct_2(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,u1_struct_0(X2))),k9_tmap_1(X1,X2),k7_tmap_1(X1,u1_struct_0(X2)))|v2_struct_0(X1)|~v1_funct_2(k9_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))|~v1_funct_1(k9_tmap_1(X1,X2))|~m1_subset_1(k9_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_111])]), c_0_18])).
cnf(c_0_121, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_112, c_0_113])).
cnf(c_0_122, negated_conjecture, (k9_tmap_1(esk1_0,esk1_0)=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk1_0))|~m1_pre_topc(esk1_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_25]), c_0_115]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_123, negated_conjecture, (esk4_3(esk1_0,esk2_0,k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))=u1_struct_0(esk2_0)|k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_116, c_0_86])])).
cnf(c_0_124, negated_conjecture, (u1_struct_0(k6_tmap_1(esk1_0,u1_struct_0(esk2_0)))=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_119])])).
cnf(c_0_125, plain, (r1_funct_2(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)),u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)),k9_tmap_1(X1,X2),k7_tmap_1(X1,u1_struct_0(X2)))|v2_struct_0(X1)|~v1_funct_2(k9_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))|~v1_funct_1(k9_tmap_1(X1,X2))|~m1_subset_1(k9_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_120, c_0_25])).
cnf(c_0_126, negated_conjecture, (k9_tmap_1(esk1_0,esk1_0)=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|~m1_pre_topc(esk1_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_29])]), c_0_31])).
cnf(c_0_127, negated_conjecture, (k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|~r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_123]), c_0_32]), c_0_124]), c_0_71]), c_0_32]), c_0_94]), c_0_86]), c_0_32]), c_0_95]), c_0_28]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_128, negated_conjecture, (r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))|~m1_pre_topc(esk1_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_115]), c_0_115]), c_0_115]), c_0_94]), c_0_86]), c_0_115]), c_0_95]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_129, negated_conjecture, (k9_tmap_1(esk1_0,esk2_0)=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|~m1_pre_topc(esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_127, c_0_128])).
cnf(c_0_130, negated_conjecture, (~r1_tmap_1(esk2_0,k8_tmap_1(esk1_0,esk2_0),k2_tmap_1(esk1_0,k8_tmap_1(esk1_0,esk2_0),k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),esk2_0),esk3_0)|~m1_pre_topc(esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_107, c_0_129])).
cnf(c_0_131, negated_conjecture, (~m1_pre_topc(esk1_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_118]), c_0_119])])).
cnf(c_0_132, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_36]), c_0_29])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 133
# Proof object clause steps            : 106
# Proof object formula steps           : 27
# Proof object conjectures             : 73
# Proof object clause conjectures      : 70
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 69
# Proof object simplifying inferences  : 226
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 3998
# ...of these trivial                  : 36
# ...subsumed                          : 2954
# ...remaining for further processing  : 1008
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 216
# Backward-rewritten                   : 141
# Generated clauses                    : 29549
# ...of the previous two non-trivial   : 28410
# Contextual simplify-reflections      : 141
# Paramodulations                      : 29545
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 618
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 602
# Current number of unprocessed clauses: 23945
# ...number of literals in the above   : 288060
# Current number of archived formulas  : 0
# Current number of archived clauses   : 387
# Clause-clause subsumption calls (NU) : 292798
# Rec. Clause-clause subsumption calls : 19027
# Non-unit clause-clause subsumptions  : 3286
# Unit Clause-clause subsumption calls : 1908
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 29
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2043145

# -------------------------------------------------
# User time                : 0.763 s
# System time              : 0.013 s
# Total time               : 0.777 s
# Maximum resident set size: 35296 pages
