# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/223363d4dfb54f94894bc2c1be23c6ad.E_input', d4_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/223363d4dfb54f94894bc2c1be23c6ad.E_input', d3_tarski)).
fof(t123_zfmisc_1, axiom, ![X1, X2, X3, X4]:k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)), file('data/MPTP/proofs/223363d4dfb54f94894bc2c1be23c6ad.E_input', t123_zfmisc_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/223363d4dfb54f94894bc2c1be23c6ad.E_input', idempotence_k3_xboole_0)).
fof(t134_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4)))), file('data/MPTP/proofs/223363d4dfb54f94894bc2c1be23c6ad.E_input', t134_zfmisc_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/223363d4dfb54f94894bc2c1be23c6ad.E_input', d5_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/223363d4dfb54f94894bc2c1be23c6ad.E_input', commutativity_k3_xboole_0)).
fof(t4_boole, axiom, ![X1]:k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0, file('data/MPTP/proofs/223363d4dfb54f94894bc2c1be23c6ad.E_input', t4_boole)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/223363d4dfb54f94894bc2c1be23c6ad.E_input', d1_tarski)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/223363d4dfb54f94894bc2c1be23c6ad.E_input', t28_xboole_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/223363d4dfb54f94894bc2c1be23c6ad.E_input', t113_zfmisc_1)).
fof(t6_xboole_0, axiom, ![X1, X2]:~((r2_xboole_0(X1,X2)&![X3]:~((r2_hidden(X3,X2)&~(r2_hidden(X3,X1)))))), file('data/MPTP/proofs/223363d4dfb54f94894bc2c1be23c6ad.E_input', t6_xboole_0)).
fof(d8_xboole_0, axiom, ![X1, X2]:(r2_xboole_0(X1,X2)<=>(r1_tarski(X1,X2)&~(X1=X2))), file('data/MPTP/proofs/223363d4dfb54f94894bc2c1be23c6ad.E_input', d8_xboole_0)).
fof(c_0_13, plain, ![X5, X6, X7, X8, X9, X10, X11, X12]:((((r2_hidden(X8,X5)|~r2_hidden(X8,X7)|X7!=k3_xboole_0(X5,X6))&(r2_hidden(X8,X6)|~r2_hidden(X8,X7)|X7!=k3_xboole_0(X5,X6)))&(~r2_hidden(X9,X5)|~r2_hidden(X9,X6)|r2_hidden(X9,X7)|X7!=k3_xboole_0(X5,X6)))&((~r2_hidden(esk1_3(X10,X11,X12),X12)|(~r2_hidden(esk1_3(X10,X11,X12),X10)|~r2_hidden(esk1_3(X10,X11,X12),X11))|X12=k3_xboole_0(X10,X11))&((r2_hidden(esk1_3(X10,X11,X12),X10)|r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k3_xboole_0(X10,X11))&(r2_hidden(esk1_3(X10,X11,X12),X11)|r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k3_xboole_0(X10,X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_14, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_15, plain, ![X14, X15, X16, X17, X18]:((~r1_tarski(X14,X15)|(~r2_hidden(X16,X14)|r2_hidden(X16,X15)))&((r2_hidden(esk2_2(X17,X18),X17)|r1_tarski(X17,X18))&(~r2_hidden(esk2_2(X17,X18),X18)|r1_tarski(X17,X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_16, plain, ![X20, X21, X22, X23]:k2_zfmisc_1(k3_xboole_0(X20,X21),k3_xboole_0(X22,X23))=k3_xboole_0(k2_zfmisc_1(X20,X22),k2_zfmisc_1(X21,X23)), inference(variable_rename,[status(thm)],[t123_zfmisc_1])).
fof(c_0_17, plain, ![X24]:k3_xboole_0(X24,X24)=X24, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
fof(c_0_18, negated_conjecture, ~(![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4))))), inference(assume_negation,[status(cth)],[t134_zfmisc_1])).
fof(c_0_19, plain, ![X29, X30, X31, X32, X33, X34, X35, X36]:((((r2_hidden(X32,X29)|~r2_hidden(X32,X31)|X31!=k4_xboole_0(X29,X30))&(~r2_hidden(X32,X30)|~r2_hidden(X32,X31)|X31!=k4_xboole_0(X29,X30)))&(~r2_hidden(X33,X29)|r2_hidden(X33,X30)|r2_hidden(X33,X31)|X31!=k4_xboole_0(X29,X30)))&((~r2_hidden(esk7_3(X34,X35,X36),X36)|(~r2_hidden(esk7_3(X34,X35,X36),X34)|r2_hidden(esk7_3(X34,X35,X36),X35))|X36=k4_xboole_0(X34,X35))&((r2_hidden(esk7_3(X34,X35,X36),X34)|r2_hidden(esk7_3(X34,X35,X36),X36)|X36=k4_xboole_0(X34,X35))&(~r2_hidden(esk7_3(X34,X35,X36),X35)|r2_hidden(esk7_3(X34,X35,X36),X36)|X36=k4_xboole_0(X34,X35))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_20, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_24, negated_conjecture, (k2_zfmisc_1(esk3_0,esk4_0)=k2_zfmisc_1(esk5_0,esk6_0)&(esk3_0!=k1_xboole_0&(esk4_0!=k1_xboole_0&(esk3_0!=esk5_0|esk4_0!=esk6_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).
fof(c_0_25, plain, ![X38, X39]:k3_xboole_0(X38,X39)=k3_xboole_0(X39,X38), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_26, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_27, plain, ![X40]:k4_xboole_0(k1_xboole_0,X40)=k1_xboole_0, inference(variable_rename,[status(thm)],[t4_boole])).
cnf(c_0_28, plain, (r1_tarski(X1,X2)|~r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_29, plain, (r1_tarski(k3_xboole_0(X1,X2),X3)|r2_hidden(esk2_2(k3_xboole_0(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_30, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))=k2_zfmisc_1(k3_xboole_0(X1,X3),X2)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_31, negated_conjecture, (k2_zfmisc_1(esk3_0,esk4_0)=k2_zfmisc_1(esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_35, plain, ![X41, X42, X43, X44, X45, X46]:(((~r2_hidden(X43,X42)|X43=X41|X42!=k1_tarski(X41))&(X44!=X41|r2_hidden(X44,X42)|X42!=k1_tarski(X41)))&((~r2_hidden(esk8_2(X45,X46),X46)|esk8_2(X45,X46)!=X45|X46=k1_tarski(X45))&(r2_hidden(esk8_2(X45,X46),X46)|esk8_2(X45,X46)=X45|X46=k1_tarski(X45)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_36, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_37, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(X1,esk4_0),k2_zfmisc_1(esk5_0,esk6_0))=k2_zfmisc_1(k3_xboole_0(X1,esk3_0),esk4_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_38, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X4),k2_zfmisc_1(X3,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_32]), c_0_22])).
cnf(c_0_39, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_40, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))=k2_zfmisc_1(X1,k3_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_42, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_36, c_0_32])).
cnf(c_0_43, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(X1,esk6_0),k2_zfmisc_1(esk5_0,esk4_0))=k2_zfmisc_1(k3_xboole_0(X1,esk3_0),esk4_0)), inference(rw,[status(thm)],[c_0_37, c_0_38])).
fof(c_0_44, plain, ![X48, X49]:(~r1_tarski(X48,X49)|k3_xboole_0(X48,X49)=X48), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_45, plain, (r1_tarski(k1_xboole_0,X1)|~r2_hidden(esk2_2(k1_xboole_0,X1),X2)), inference(spm,[status(thm)],[c_0_39, c_0_21])).
cnf(c_0_46, plain, (r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_40])])).
cnf(c_0_47, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk5_0,esk6_0),k2_zfmisc_1(esk3_0,X1))=k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk4_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_31]), c_0_32])).
cnf(c_0_48, negated_conjecture, (r1_tarski(k2_zfmisc_1(k3_xboole_0(X1,esk3_0),esk4_0),k2_zfmisc_1(X1,esk6_0))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_49, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_50, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_51, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_52, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk5_0,X1),k2_zfmisc_1(esk3_0,esk6_0))=k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk4_0))), inference(spm,[status(thm)],[c_0_47, c_0_38])).
cnf(c_0_53, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_54, negated_conjecture, (r1_tarski(k2_zfmisc_1(k3_xboole_0(esk3_0,X1),esk4_0),k2_zfmisc_1(X1,esk6_0))), inference(spm,[status(thm)],[c_0_48, c_0_32])).
cnf(c_0_55, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_36]), c_0_32])).
cnf(c_0_56, plain, (~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_39])).
cnf(c_0_57, plain, (r2_hidden(esk1_3(X1,X2,X3),X2)|r2_hidden(esk1_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_58, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk4_0)),k2_zfmisc_1(esk5_0,X1))), inference(spm,[status(thm)],[c_0_42, c_0_52])).
cnf(c_0_59, plain, (X1=X2|~r2_hidden(X1,k1_tarski(X2))), inference(er,[status(thm)],[c_0_53])).
fof(c_0_60, plain, ![X50, X51]:((k2_zfmisc_1(X50,X51)!=k1_xboole_0|(X50=k1_xboole_0|X51=k1_xboole_0))&((X50!=k1_xboole_0|k2_zfmisc_1(X50,X51)=k1_xboole_0)&(X51!=k1_xboole_0|k2_zfmisc_1(X50,X51)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_61, negated_conjecture, (r1_tarski(k2_zfmisc_1(k3_xboole_0(X1,esk3_0),esk4_0),k2_zfmisc_1(k3_xboole_0(X1,esk3_0),esk6_0))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_62, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|r2_hidden(esk1_3(X1,X2,k1_xboole_0),X2)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_63, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,k3_xboole_0(esk4_0,X1)),k2_zfmisc_1(esk5_0,X1))), inference(spm,[status(thm)],[c_0_58, c_0_32])).
cnf(c_0_64, plain, (esk2_2(k1_tarski(X1),X2)=X1|r1_tarski(k1_tarski(X1),X2)), inference(spm,[status(thm)],[c_0_59, c_0_21])).
cnf(c_0_65, plain, (X1=k1_xboole_0|X2=k1_xboole_0|k2_zfmisc_1(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_66, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(X1,esk3_0),esk4_0)=k2_zfmisc_1(k3_xboole_0(X1,esk5_0),esk6_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_61]), c_0_41]), c_0_32]), c_0_22]), c_0_31]), c_0_30])).
cnf(c_0_67, negated_conjecture, (esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_68, plain, (r2_hidden(esk1_3(X1,X2,X3),X1)|r2_hidden(esk1_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_69, plain, (esk1_3(X1,k1_tarski(X2),k1_xboole_0)=X2|k3_xboole_0(X1,k1_tarski(X2))=k1_xboole_0), inference(spm,[status(thm)],[c_0_59, c_0_62])).
cnf(c_0_70, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk4_0)),k2_zfmisc_1(esk5_0,k3_xboole_0(X1,esk4_0)))), inference(spm,[status(thm)],[c_0_63, c_0_55])).
cnf(c_0_71, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_28, c_0_64])).
cnf(c_0_72, negated_conjecture, (k3_xboole_0(X1,esk3_0)=k1_xboole_0|k2_zfmisc_1(k3_xboole_0(X1,esk5_0),esk6_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67])).
cnf(c_0_73, plain, (k3_xboole_0(X1,k1_tarski(X2))=k1_xboole_0|r2_hidden(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_56])).
cnf(c_0_74, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_75, negated_conjecture, (k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk4_0))=k2_zfmisc_1(esk5_0,k3_xboole_0(X1,esk6_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_70]), c_0_30]), c_0_32]), c_0_22]), c_0_31]), c_0_41])).
cnf(c_0_76, plain, (r1_tarski(k2_zfmisc_1(X1,k3_xboole_0(X2,X3)),k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_42, c_0_41])).
cnf(c_0_77, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk5_0,esk3_0),esk6_0)=k2_zfmisc_1(esk3_0,k3_xboole_0(esk6_0,esk4_0))), inference(spm,[status(thm)],[c_0_30, c_0_47])).
cnf(c_0_78, plain, (r1_tarski(k1_tarski(esk2_2(X1,X2)),X1)|r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_71, c_0_21])).
cnf(c_0_79, negated_conjecture, (k3_xboole_0(X1,esk3_0)=k1_xboole_0|k2_zfmisc_1(k3_xboole_0(esk5_0,X1),esk6_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_72, c_0_32])).
cnf(c_0_80, plain, (k3_xboole_0(X1,k1_tarski(esk2_2(X2,X1)))=k1_xboole_0|r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_28, c_0_73])).
cnf(c_0_81, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_74])).
cnf(c_0_82, negated_conjecture, (k2_zfmisc_1(esk3_0,k3_xboole_0(esk4_0,X1))=k2_zfmisc_1(esk5_0,k3_xboole_0(X1,esk6_0))), inference(spm,[status(thm)],[c_0_75, c_0_32])).
cnf(c_0_83, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_84, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk4_0)),k2_zfmisc_1(esk3_0,k3_xboole_0(esk6_0,esk4_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_22]), c_0_47])).
cnf(c_0_85, plain, (k3_xboole_0(X1,k1_tarski(esk2_2(X1,X2)))=k1_tarski(esk2_2(X1,X2))|r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_78]), c_0_32])).
cnf(c_0_86, negated_conjecture, (k3_xboole_0(esk3_0,k1_tarski(esk2_2(X1,esk5_0)))=k1_xboole_0|r1_tarski(X1,esk5_0)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_81])]), c_0_32])).
cnf(c_0_87, negated_conjecture, (k3_xboole_0(esk4_0,X1)=k1_xboole_0|k2_zfmisc_1(esk5_0,k3_xboole_0(X1,esk6_0))!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_82]), c_0_83])).
cnf(c_0_88, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_89, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk5_0,esk6_0),k2_zfmisc_1(esk3_0,k3_xboole_0(esk6_0,esk4_0)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_23]), c_0_31])).
cnf(c_0_90, negated_conjecture, (k1_tarski(esk2_2(esk3_0,esk5_0))=k1_xboole_0|r1_tarski(esk3_0,esk5_0)), inference(spm,[status(thm)],[c_0_85, c_0_86])).
fof(c_0_91, plain, ![X52, X53]:((r2_hidden(esk9_2(X52,X53),X53)|~r2_xboole_0(X52,X53))&(~r2_hidden(esk9_2(X52,X53),X52)|~r2_xboole_0(X52,X53))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_xboole_0])])])])])).
cnf(c_0_92, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk5_0,esk6_0),k2_zfmisc_1(X1,esk4_0))=k2_zfmisc_1(k3_xboole_0(esk3_0,X1),esk4_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_93, negated_conjecture, (k3_xboole_0(esk4_0,X1)=k1_xboole_0|k2_zfmisc_1(esk5_0,k3_xboole_0(esk6_0,X1))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_87, c_0_32])).
cnf(c_0_94, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_88])).
cnf(c_0_95, negated_conjecture, (k2_zfmisc_1(esk3_0,k3_xboole_0(esk6_0,esk4_0))=k2_zfmisc_1(esk5_0,esk6_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_89]), c_0_47]), c_0_32]), c_0_55])).
cnf(c_0_96, negated_conjecture, (r1_tarski(esk3_0,esk5_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_90]), c_0_56])).
cnf(c_0_97, plain, (r2_hidden(esk9_2(X1,X2),X2)|~r2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_98, plain, (r1_tarski(k2_zfmisc_1(X1,k3_xboole_0(X2,X3)),k2_zfmisc_1(X1,X3))), inference(spm,[status(thm)],[c_0_36, c_0_41])).
cnf(c_0_99, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk5_0,esk3_0),esk4_0)=k2_zfmisc_1(esk5_0,k3_xboole_0(esk6_0,esk4_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_41]), c_0_32])).
cnf(c_0_100, negated_conjecture, (k3_xboole_0(esk4_0,k1_tarski(esk2_2(X1,esk6_0)))=k1_xboole_0|r1_tarski(X1,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_80]), c_0_94])])).
cnf(c_0_101, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk5_0,esk3_0),esk6_0)=k2_zfmisc_1(esk5_0,esk6_0)), inference(rw,[status(thm)],[c_0_77, c_0_95])).
cnf(c_0_102, negated_conjecture, (k3_xboole_0(esk5_0,esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_96]), c_0_32])).
cnf(c_0_103, plain, (~r2_xboole_0(X1,k1_xboole_0)|~r2_hidden(esk9_2(X1,k1_xboole_0),X2)), inference(spm,[status(thm)],[c_0_39, c_0_97])).
fof(c_0_104, plain, ![X55, X56]:(((r1_tarski(X55,X56)|~r2_xboole_0(X55,X56))&(X55!=X56|~r2_xboole_0(X55,X56)))&(~r1_tarski(X55,X56)|X55=X56|r2_xboole_0(X55,X56))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_xboole_0])])])])).
cnf(c_0_105, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_20, c_0_32])).
cnf(c_0_106, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk5_0,k3_xboole_0(X1,esk6_0)),k2_zfmisc_1(esk5_0,k3_xboole_0(esk6_0,esk4_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_22]), c_0_31]), c_0_41])).
cnf(c_0_107, negated_conjecture, (k1_tarski(esk2_2(esk4_0,esk6_0))=k1_xboole_0|r1_tarski(esk4_0,esk6_0)), inference(spm,[status(thm)],[c_0_85, c_0_100])).
cnf(c_0_108, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_109, negated_conjecture, (k2_zfmisc_1(esk3_0,esk6_0)=k2_zfmisc_1(esk5_0,esk6_0)), inference(rw,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_110, plain, (~r2_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_103, c_0_46])).
cnf(c_0_111, plain, (X1=X2|r2_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_112, plain, (r1_tarski(k3_xboole_0(X1,k4_xboole_0(X2,X3)),X4)|~r2_hidden(esk2_2(k3_xboole_0(X1,k4_xboole_0(X2,X3)),X4),X3)), inference(spm,[status(thm)],[c_0_33, c_0_29])).
cnf(c_0_113, plain, (r1_tarski(k3_xboole_0(X1,X2),X3)|r2_hidden(esk2_2(k3_xboole_0(X1,X2),X3),X1)), inference(spm,[status(thm)],[c_0_105, c_0_21])).
cnf(c_0_114, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk5_0,esk6_0),k2_zfmisc_1(esk5_0,k3_xboole_0(esk6_0,esk4_0)))), inference(spm,[status(thm)],[c_0_106, c_0_23])).
cnf(c_0_115, plain, (k3_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_42]), c_0_32])).
cnf(c_0_116, negated_conjecture, (r1_tarski(esk4_0,esk6_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_107]), c_0_56])).
cnf(c_0_117, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=k3_xboole_0(k2_zfmisc_1(X3,X2),k2_zfmisc_1(X1,X4))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_32]), c_0_22])).
cnf(c_0_118, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_108])).
cnf(c_0_119, negated_conjecture, (k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk6_0))=k2_zfmisc_1(esk5_0,k3_xboole_0(X1,esk6_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_109]), c_0_32]), c_0_47]), c_0_75])).
cnf(c_0_120, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_110, c_0_111])).
cnf(c_0_121, plain, (r1_tarski(k3_xboole_0(X1,k4_xboole_0(X2,X1)),X3)), inference(spm,[status(thm)],[c_0_112, c_0_113])).
cnf(c_0_122, negated_conjecture, (k2_zfmisc_1(esk5_0,k3_xboole_0(esk6_0,esk4_0))=k2_zfmisc_1(esk5_0,esk6_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_114]), c_0_41]), c_0_115])).
cnf(c_0_123, negated_conjecture, (k3_xboole_0(esk6_0,esk4_0)=esk4_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_116]), c_0_32])).
cnf(c_0_124, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(X1,esk5_0),esk6_0)=k2_zfmisc_1(k3_xboole_0(esk3_0,X1),esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_117]), c_0_43]), c_0_66])).
cnf(c_0_125, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|r2_hidden(esk2_2(k4_xboole_0(X1,X2),X3),X1)), inference(spm,[status(thm)],[c_0_118, c_0_21])).
cnf(c_0_126, negated_conjecture, (k3_xboole_0(X1,esk6_0)=k1_xboole_0|k2_zfmisc_1(esk5_0,k3_xboole_0(X1,esk6_0))!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_119]), c_0_83])).
cnf(c_0_127, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_128, negated_conjecture, (k2_zfmisc_1(esk5_0,esk4_0)=k2_zfmisc_1(esk5_0,esk6_0)), inference(rw,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_129, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk3_0,X1),esk4_0)=k2_zfmisc_1(k3_xboole_0(esk5_0,X1),esk6_0)), inference(spm,[status(thm)],[c_0_124, c_0_32])).
cnf(c_0_130, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_28, c_0_125])).
cnf(c_0_131, negated_conjecture, (k3_xboole_0(esk6_0,X1)=k1_xboole_0|k2_zfmisc_1(esk5_0,k3_xboole_0(esk6_0,X1))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_126, c_0_32])).
cnf(c_0_132, negated_conjecture, (k2_zfmisc_1(esk5_0,k3_xboole_0(esk6_0,k4_xboole_0(X1,esk4_0)))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_127]), c_0_94]), c_0_32])).
cnf(c_0_133, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk5_0,X1),esk4_0)=k2_zfmisc_1(k3_xboole_0(esk5_0,X1),esk6_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_128]), c_0_92]), c_0_129])).
cnf(c_0_134, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_135, plain, (k3_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_130]), c_0_32])).
cnf(c_0_136, negated_conjecture, (k3_xboole_0(esk6_0,k4_xboole_0(X1,esk4_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_131, c_0_132])).
cnf(c_0_137, negated_conjecture, (k3_xboole_0(esk5_0,X1)=k1_xboole_0|k2_zfmisc_1(k3_xboole_0(esk5_0,X1),esk6_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_133]), c_0_67])).
cnf(c_0_138, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk5_0,k4_xboole_0(X1,esk3_0)),esk6_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_127]), c_0_81])).
cnf(c_0_139, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_134])).
cnf(c_0_140, negated_conjecture, (k4_xboole_0(esk6_0,esk4_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_141, negated_conjecture, (k3_xboole_0(esk5_0,k4_xboole_0(X1,esk3_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_137, c_0_138])).
cnf(c_0_142, plain, (~r2_hidden(esk9_2(X1,X2),X1)|~r2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_143, negated_conjecture, (r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk6_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140]), c_0_56])).
cnf(c_0_144, negated_conjecture, (k4_xboole_0(esk5_0,esk3_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_135, c_0_141])).
cnf(c_0_145, negated_conjecture, (~r2_xboole_0(esk4_0,X1)|~r2_hidden(esk9_2(esk4_0,X1),esk6_0)), inference(spm,[status(thm)],[c_0_142, c_0_143])).
cnf(c_0_146, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk5_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_144]), c_0_56])).
cnf(c_0_147, negated_conjecture, (~r2_xboole_0(esk4_0,esk6_0)), inference(spm,[status(thm)],[c_0_145, c_0_97])).
cnf(c_0_148, negated_conjecture, (~r2_xboole_0(esk3_0,X1)|~r2_hidden(esk9_2(esk3_0,X1),esk5_0)), inference(spm,[status(thm)],[c_0_142, c_0_146])).
cnf(c_0_149, negated_conjecture, (esk3_0!=esk5_0|esk4_0!=esk6_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_150, negated_conjecture, (esk4_0=esk6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_111]), c_0_116])])).
cnf(c_0_151, negated_conjecture, (~r2_xboole_0(esk3_0,esk5_0)), inference(spm,[status(thm)],[c_0_148, c_0_97])).
cnf(c_0_152, negated_conjecture, (esk3_0!=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_149, c_0_150])])).
cnf(c_0_153, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_111]), c_0_96])]), c_0_152]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 154
# Proof object clause steps            : 127
# Proof object formula steps           : 27
# Proof object conjectures             : 66
# Proof object clause conjectures      : 63
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 88
# Proof object simplifying inferences  : 76
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 36
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 36
# Processed clauses                    : 1441
# ...of these trivial                  : 259
# ...subsumed                          : 591
# ...remaining for further processing  : 591
# Other redundant clauses eliminated   : 15
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 258
# Generated clauses                    : 12188
# ...of the previous two non-trivial   : 8039
# Contextual simplify-reflections      : 1
# Paramodulations                      : 12166
# Factorizations                       : 8
# Equation resolutions                 : 15
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 321
#    Positive orientable unit clauses  : 145
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 21
#    Non-unit-clauses                  : 152
# Current number of unprocessed clauses: 6487
# ...number of literals in the above   : 10577
# Current number of archived formulas  : 0
# Current number of archived clauses   : 259
# Clause-clause subsumption calls (NU) : 4288
# Rec. Clause-clause subsumption calls : 3478
# Non-unit clause-clause subsumptions  : 222
# Unit Clause-clause subsumption calls : 1784
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 409
# BW rewrite match successes           : 110
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 147365

# -------------------------------------------------
# User time                : 0.098 s
# System time              : 0.000 s
# Total time               : 0.098 s
# Maximum resident set size: 28648 pages
