# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t31_subset_1, conjecture, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(r1_tarski(X2,X3)<=>r1_tarski(k3_subset_1(X1,X3),k3_subset_1(X1,X2))))), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', t31_subset_1)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', d1_zfmisc_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', d2_subset_1)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', fc1_subset_1)).
fof(t69_xboole_1, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>~((r1_tarski(X2,X1)&r1_xboole_0(X2,X1)))), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', t69_xboole_1)).
fof(t85_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_xboole_0(X1,k4_xboole_0(X3,X2))), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', t85_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', d10_xboole_0)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', t1_xboole_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', t36_xboole_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', d1_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', d3_tarski)).
fof(t33_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(k4_xboole_0(X1,X3),k4_xboole_0(X2,X3))), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', t33_xboole_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', t2_xboole_1)).
fof(t34_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(k4_xboole_0(X3,X2),k4_xboole_0(X3,X1))), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', t34_xboole_1)).
fof(t106_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k4_xboole_0(X2,X3))=>(r1_tarski(X1,X2)&r1_xboole_0(X1,X3))), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', t106_xboole_1)).
fof(d5_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', d5_subset_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', t7_xboole_1)).
fof(t43_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3)), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', t43_xboole_1)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', t1_boole)).
fof(t86_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_xboole_0(X1,X3))=>r1_tarski(X1,k4_xboole_0(X2,X3))), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', t86_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', commutativity_k2_xboole_0)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/224a6e44e0384649afdfee4f312ffbc4.E_input', t39_xboole_1)).
fof(c_0_22, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(r1_tarski(X2,X3)<=>r1_tarski(k3_subset_1(X1,X3),k3_subset_1(X1,X2)))))), inference(assume_negation,[status(cth)],[t31_subset_1])).
fof(c_0_23, plain, ![X7, X8, X9, X10, X11, X12]:(((~r2_hidden(X9,X8)|r1_tarski(X9,X7)|X8!=k1_zfmisc_1(X7))&(~r1_tarski(X10,X7)|r2_hidden(X10,X8)|X8!=k1_zfmisc_1(X7)))&((~r2_hidden(esk4_2(X11,X12),X12)|~r1_tarski(esk4_2(X11,X12),X11)|X12=k1_zfmisc_1(X11))&(r2_hidden(esk4_2(X11,X12),X12)|r1_tarski(esk4_2(X11,X12),X11)|X12=k1_zfmisc_1(X11)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
fof(c_0_24, plain, ![X14, X15]:(((~m1_subset_1(X15,X14)|r2_hidden(X15,X14)|v1_xboole_0(X14))&(~r2_hidden(X15,X14)|m1_subset_1(X15,X14)|v1_xboole_0(X14)))&((~m1_subset_1(X15,X14)|v1_xboole_0(X15)|~v1_xboole_0(X14))&(~v1_xboole_0(X15)|m1_subset_1(X15,X14)|~v1_xboole_0(X14)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_25, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))&((~r1_tarski(esk2_0,esk3_0)|~r1_tarski(k3_subset_1(esk1_0,esk3_0),k3_subset_1(esk1_0,esk2_0)))&(r1_tarski(esk2_0,esk3_0)|r1_tarski(k3_subset_1(esk1_0,esk3_0),k3_subset_1(esk1_0,esk2_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])).
fof(c_0_26, plain, ![X16]:~v1_xboole_0(k1_zfmisc_1(X16)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
cnf(c_0_27, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_31, plain, ![X17, X18]:(v1_xboole_0(X18)|(~r1_tarski(X18,X17)|~r1_xboole_0(X18,X17))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t69_xboole_1])])])).
fof(c_0_32, plain, ![X19, X20, X21]:(~r1_tarski(X19,X20)|r1_xboole_0(X19,k4_xboole_0(X21,X20))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t85_xboole_1])])).
fof(c_0_33, plain, ![X22, X23]:(((r1_tarski(X22,X23)|X22!=X23)&(r1_tarski(X23,X22)|X22!=X23))&(~r1_tarski(X22,X23)|~r1_tarski(X23,X22)|X22=X23)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_34, plain, ![X24, X25, X26]:(~r1_tarski(X24,X25)|~r1_tarski(X25,X26)|r1_tarski(X24,X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
cnf(c_0_35, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_36, negated_conjecture, (r2_hidden(esk2_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])).
cnf(c_0_37, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_38, plain, (v1_xboole_0(X1)|~r1_tarski(X1,X2)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_39, plain, (r1_xboole_0(X1,k4_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_41, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42, negated_conjecture, (r1_tarski(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
fof(c_0_43, plain, ![X27, X28]:r1_tarski(k4_xboole_0(X27,X28),X27), inference(variable_rename,[status(thm)],[t36_xboole_1])).
cnf(c_0_44, negated_conjecture, (r2_hidden(esk3_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_37]), c_0_30])).
fof(c_0_45, plain, ![X29, X30, X31]:((~v1_xboole_0(X29)|~r2_hidden(X30,X29))&(r2_hidden(esk5_1(X31),X31)|v1_xboole_0(X31))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_46, plain, ![X33, X34, X35, X36, X37]:((~r1_tarski(X33,X34)|(~r2_hidden(X35,X33)|r2_hidden(X35,X34)))&((r2_hidden(esk6_2(X36,X37),X36)|r1_tarski(X36,X37))&(~r2_hidden(esk6_2(X36,X37),X37)|r1_tarski(X36,X37)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_47, plain, (v1_xboole_0(X1)|~r1_tarski(X1,k4_xboole_0(X2,X3))|~r1_tarski(X1,X3)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_48, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_49, negated_conjecture, (r1_tarski(X1,esk1_0)|~r1_tarski(X1,esk2_0)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_50, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_51, plain, ![X39, X40, X41]:(~r1_tarski(X39,X40)|r1_tarski(k4_xboole_0(X39,X41),k4_xboole_0(X40,X41))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t33_xboole_1])])).
cnf(c_0_52, negated_conjecture, (r1_tarski(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_35, c_0_44])).
fof(c_0_53, plain, ![X42]:r1_tarski(k1_xboole_0,X42), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_54, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_55, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_56, plain, (v1_xboole_0(k4_xboole_0(X1,X2))|~r1_tarski(k4_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_57, negated_conjecture, (r1_tarski(k4_xboole_0(esk2_0,X1),esk1_0)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_58, plain, (r1_tarski(k4_xboole_0(X1,X3),k4_xboole_0(X2,X3))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_59, negated_conjecture, (r1_tarski(X1,esk1_0)|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_41, c_0_52])).
fof(c_0_60, plain, ![X43, X44, X45]:(~r1_tarski(X43,X44)|r1_tarski(k4_xboole_0(X45,X44),k4_xboole_0(X45,X43))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t34_xboole_1])])).
cnf(c_0_61, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_62, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_63, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_64, negated_conjecture, (v1_xboole_0(k4_xboole_0(esk2_0,esk1_0))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
fof(c_0_65, plain, ![X46, X47, X48]:((r1_tarski(X46,X47)|~r1_tarski(X46,k4_xboole_0(X47,X48)))&(r1_xboole_0(X46,X48)|~r1_tarski(X46,k4_xboole_0(X47,X48)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t106_xboole_1])])])).
fof(c_0_66, plain, ![X49, X50]:(~m1_subset_1(X50,k1_zfmisc_1(X49))|k3_subset_1(X49,X50)=k4_xboole_0(X49,X50)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
cnf(c_0_67, plain, (v1_xboole_0(k4_xboole_0(X1,X2))|~r1_tarski(k4_xboole_0(X1,X2),X2)|~r1_tarski(X1,X3)), inference(spm,[status(thm)],[c_0_47, c_0_58])).
cnf(c_0_68, negated_conjecture, (r1_tarski(k4_xboole_0(esk3_0,X1),esk1_0)), inference(spm,[status(thm)],[c_0_59, c_0_50])).
fof(c_0_69, plain, ![X51, X52]:r1_tarski(X51,k2_xboole_0(X51,X52)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_70, plain, ![X53, X54, X55]:(~r1_tarski(X53,k2_xboole_0(X54,X55))|r1_tarski(k4_xboole_0(X53,X54),X55)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
fof(c_0_71, plain, ![X56]:k2_xboole_0(X56,k1_xboole_0)=X56, inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_72, plain, (r1_tarski(k4_xboole_0(X3,X2),k4_xboole_0(X3,X1))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_73, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_74, negated_conjecture, (r1_tarski(k4_xboole_0(esk2_0,esk1_0),X1)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
fof(c_0_75, plain, ![X57, X58, X59]:(~r1_tarski(X57,X58)|~r1_xboole_0(X57,X59)|r1_tarski(X57,k4_xboole_0(X58,X59))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t86_xboole_1])])).
cnf(c_0_76, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k4_xboole_0(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_77, plain, (k3_subset_1(X2,X1)=k4_xboole_0(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_78, negated_conjecture, (v1_xboole_0(k4_xboole_0(esk3_0,esk1_0))|~r1_tarski(esk3_0,X1)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_79, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_80, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_81, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_71])).
fof(c_0_82, plain, ![X60, X61]:k2_xboole_0(X60,X61)=k2_xboole_0(X61,X60), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_83, plain, ![X62, X63]:k2_xboole_0(X62,k4_xboole_0(X63,X62))=k2_xboole_0(X62,X63), inference(variable_rename,[status(thm)],[t39_xboole_1])).
cnf(c_0_84, plain, (k4_xboole_0(X1,X2)=k4_xboole_0(X1,X3)|~r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(X1,X3))|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_61, c_0_72])).
cnf(c_0_85, negated_conjecture, (k4_xboole_0(esk2_0,esk1_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_86, plain, (r1_tarski(X1,k4_xboole_0(X2,X3))|~r1_tarski(X1,X2)|~r1_xboole_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_87, plain, (r1_xboole_0(k4_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_76, c_0_48])).
cnf(c_0_88, negated_conjecture, (~r1_tarski(esk2_0,esk3_0)|~r1_tarski(k3_subset_1(esk1_0,esk3_0),k3_subset_1(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_89, negated_conjecture, (k3_subset_1(esk1_0,esk3_0)=k4_xboole_0(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_77, c_0_37])).
cnf(c_0_90, negated_conjecture, (k3_subset_1(esk1_0,esk2_0)=k4_xboole_0(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_77, c_0_29])).
cnf(c_0_91, negated_conjecture, (v1_xboole_0(k4_xboole_0(esk3_0,esk1_0))), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_92, plain, (r1_tarski(k4_xboole_0(X1,X2),k1_xboole_0)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_93, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_94, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_95, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_96, negated_conjecture, (k4_xboole_0(esk2_0,X1)=k1_xboole_0|~r1_tarski(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_62])])).
cnf(c_0_97, plain, (k4_xboole_0(X1,X2)=k4_xboole_0(X3,X2)|~r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(X3,X2))|~r1_tarski(X3,X1)), inference(spm,[status(thm)],[c_0_61, c_0_58])).
cnf(c_0_98, plain, (r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(X3,X2))|~r1_tarski(k4_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_99, negated_conjecture, (r1_tarski(esk2_0,esk3_0)|r1_tarski(k3_subset_1(esk1_0,esk3_0),k3_subset_1(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_100, negated_conjecture, (~r1_tarski(esk2_0,esk3_0)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_88, c_0_89]), c_0_90]), c_0_72])).
cnf(c_0_101, negated_conjecture, (r1_tarski(k4_xboole_0(esk3_0,esk1_0),X1)), inference(spm,[status(thm)],[c_0_63, c_0_91])).
cnf(c_0_102, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_92]), c_0_62])])).
cnf(c_0_103, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_93, c_0_50])).
cnf(c_0_104, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_79, c_0_94])).
cnf(c_0_105, negated_conjecture, (k2_xboole_0(X1,esk2_0)=X1|~r1_tarski(esk1_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_81])).
cnf(c_0_106, plain, (k4_xboole_0(X1,X2)=k4_xboole_0(X3,X2)|~r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X3,X1)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_107, negated_conjecture, (r1_tarski(k4_xboole_0(esk1_0,esk3_0),k4_xboole_0(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_99, c_0_89]), c_0_90]), c_0_100])).
cnf(c_0_108, negated_conjecture, (k4_xboole_0(esk3_0,esk1_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_73, c_0_101])).
cnf(c_0_109, plain, (k2_xboole_0(X1,X2)=X1|~r1_tarski(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_102]), c_0_81])).
cnf(c_0_110, plain, (r2_hidden(esk6_2(k4_xboole_0(X1,X2),X3),X1)|r1_tarski(k4_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_103, c_0_55])).
cnf(c_0_111, negated_conjecture, (r1_tarski(esk2_0,X1)|~r1_tarski(esk1_0,X1)), inference(spm,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_112, negated_conjecture, (k4_xboole_0(k4_xboole_0(esk1_0,esk2_0),esk3_0)=k4_xboole_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_50])])).
cnf(c_0_113, negated_conjecture, (k2_xboole_0(esk1_0,esk3_0)=esk1_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_108]), c_0_81])).
cnf(c_0_114, plain, (k2_xboole_0(X1,X2)=X1|~r1_tarski(k4_xboole_0(X2,X1),X1)), inference(spm,[status(thm)],[c_0_95, c_0_109])).
cnf(c_0_115, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_54, c_0_110])).
cnf(c_0_116, negated_conjecture, (r1_tarski(k4_xboole_0(esk2_0,X1),X2)|~r1_tarski(esk1_0,k2_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_80, c_0_111])).
cnf(c_0_117, negated_conjecture, (k2_xboole_0(esk3_0,k4_xboole_0(esk1_0,esk2_0))=esk1_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_112]), c_0_95]), c_0_94]), c_0_113])).
cnf(c_0_118, plain, (k2_xboole_0(k1_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_81, c_0_94])).
cnf(c_0_119, plain, (k2_xboole_0(X1,X2)=X1|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_120, negated_conjecture, (r1_tarski(k4_xboole_0(esk2_0,esk3_0),k4_xboole_0(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_48])])).
cnf(c_0_121, plain, (k1_xboole_0=X1|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_118, c_0_119])).
cnf(c_0_122, negated_conjecture, (v1_xboole_0(k4_xboole_0(esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_120]), c_0_50])])).
cnf(c_0_123, negated_conjecture, (k4_xboole_0(esk2_0,esk3_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_121, c_0_122])).
cnf(c_0_124, negated_conjecture, (k2_xboole_0(esk2_0,esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_123]), c_0_81]), c_0_94])).
cnf(c_0_125, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_124]), c_0_100]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 126
# Proof object clause steps            : 81
# Proof object formula steps           : 45
# Proof object conjectures             : 38
# Proof object clause conjectures      : 35
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 22
# Proof object generating inferences   : 50
# Proof object simplifying inferences  : 29
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 37
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 37
# Processed clauses                    : 31852
# ...of these trivial                  : 134
# ...subsumed                          : 28259
# ...remaining for further processing  : 3459
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 291
# Backward-rewritten                   : 76
# Generated clauses                    : 358773
# ...of the previous two non-trivial   : 324673
# Contextual simplify-reflections      : 65
# Paramodulations                      : 358700
# Factorizations                       : 68
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 3051
#    Positive orientable unit clauses  : 283
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 80
#    Non-unit-clauses                  : 2687
# Current number of unprocessed clauses: 291688
# ...number of literals in the above   : 1062663
# Current number of archived formulas  : 0
# Current number of archived clauses   : 404
# Clause-clause subsumption calls (NU) : 1862208
# Rec. Clause-clause subsumption calls : 1016626
# Non-unit clause-clause subsumptions  : 12029
# Unit Clause-clause subsumption calls : 44917
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3917
# BW rewrite match successes           : 50
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6077229

# -------------------------------------------------
# User time                : 4.484 s
# System time              : 0.060 s
# Total time               : 4.544 s
# Maximum resident set size: 221100 pages
