# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t36_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>r1_lattice3(X1,k4_waybel_0(X1,X2),X3))))), file('data/MPTP/proofs/229a5b08a97d483f8c6a51565539b039.E_input', t36_waybel_0)).
fof(t9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(r1_tarski(X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X3,X4)=>r1_lattice3(X1,X2,X4))&(r2_lattice3(X1,X3,X4)=>r2_lattice3(X1,X2,X4)))))), file('data/MPTP/proofs/229a5b08a97d483f8c6a51565539b039.E_input', t9_yellow_0)).
fof(t16_waybel_0, axiom, ![X1]:((v3_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(r1_tarski(X2,k3_waybel_0(X1,X2))&r1_tarski(X2,k4_waybel_0(X1,X2))))), file('data/MPTP/proofs/229a5b08a97d483f8c6a51565539b039.E_input', t16_waybel_0)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/229a5b08a97d483f8c6a51565539b039.E_input', d8_lattice3)).
fof(d16_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k4_waybel_0(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>?[X5]:(m1_subset_1(X5,u1_struct_0(X1))&(r1_orders_2(X1,X5,X4)&r2_hidden(X5,X2))))))))), file('data/MPTP/proofs/229a5b08a97d483f8c6a51565539b039.E_input', d16_waybel_0)).
fof(t4_yellow_0, axiom, ![X1]:((v4_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)=>![X4]:((r1_lattice3(X1,X4,X3)=>r1_lattice3(X1,X4,X2))&(r2_lattice3(X1,X4,X2)=>r2_lattice3(X1,X4,X3))))))), file('data/MPTP/proofs/229a5b08a97d483f8c6a51565539b039.E_input', t4_yellow_0)).
fof(t7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2)))))))), file('data/MPTP/proofs/229a5b08a97d483f8c6a51565539b039.E_input', t7_yellow_0)).
fof(dt_k4_waybel_0, axiom, ![X1, X2]:((l1_orders_2(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k4_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/229a5b08a97d483f8c6a51565539b039.E_input', dt_k4_waybel_0)).
fof(c_0_8, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>r1_lattice3(X1,k4_waybel_0(X1,X2),X3)))))), inference(assume_negation,[status(cth)],[t36_waybel_0])).
fof(c_0_9, plain, ![X12, X13, X14, X15]:((~r1_lattice3(X12,X14,X15)|r1_lattice3(X12,X13,X15)|~m1_subset_1(X15,u1_struct_0(X12))|~r1_tarski(X13,X14)|~l1_orders_2(X12))&(~r2_lattice3(X12,X14,X15)|r2_lattice3(X12,X13,X15)|~m1_subset_1(X15,u1_struct_0(X12))|~r1_tarski(X13,X14)|~l1_orders_2(X12))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_yellow_0])])])])).
fof(c_0_10, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&((~r1_lattice3(esk1_0,esk2_0,esk3_0)|~r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0))&(r1_lattice3(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])).
fof(c_0_11, plain, ![X32, X33]:((r1_tarski(X33,k3_waybel_0(X32,X33))|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32)))|(~v3_orders_2(X32)|~l1_orders_2(X32)))&(r1_tarski(X33,k4_waybel_0(X32,X33))|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32)))|(~v3_orders_2(X32)|~l1_orders_2(X32)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_waybel_0])])])])).
cnf(c_0_12, plain, (r1_lattice3(X1,X4,X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_tarski(X4,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (r1_tarski(X1,k4_waybel_0(X2,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v3_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_18, plain, ![X25, X26, X27, X28]:((~r1_lattice3(X25,X26,X27)|(~m1_subset_1(X28,u1_struct_0(X25))|(~r2_hidden(X28,X26)|r1_orders_2(X25,X27,X28)))|~m1_subset_1(X27,u1_struct_0(X25))|~l1_orders_2(X25))&((m1_subset_1(esk7_3(X25,X26,X27),u1_struct_0(X25))|r1_lattice3(X25,X26,X27)|~m1_subset_1(X27,u1_struct_0(X25))|~l1_orders_2(X25))&((r2_hidden(esk7_3(X25,X26,X27),X26)|r1_lattice3(X25,X26,X27)|~m1_subset_1(X27,u1_struct_0(X25))|~l1_orders_2(X25))&(~r1_orders_2(X25,X27,esk7_3(X25,X26,X27))|r1_lattice3(X25,X26,X27)|~m1_subset_1(X27,u1_struct_0(X25))|~l1_orders_2(X25))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
cnf(c_0_19, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,X1,esk3_0)|~r1_tarski(X1,k4_waybel_0(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_15])])).
cnf(c_0_20, negated_conjecture, (r1_tarski(X1,k4_waybel_0(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_15])])).
cnf(c_0_21, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_22, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])])).
fof(c_0_24, plain, ![X16, X17, X18, X19, X21, X23]:((((m1_subset_1(esk4_4(X16,X17,X18,X19),u1_struct_0(X16))|~r2_hidden(X19,X18)|~m1_subset_1(X19,u1_struct_0(X16))|X18!=k4_waybel_0(X16,X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X16)))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~l1_orders_2(X16))&((r1_orders_2(X16,esk4_4(X16,X17,X18,X19),X19)|~r2_hidden(X19,X18)|~m1_subset_1(X19,u1_struct_0(X16))|X18!=k4_waybel_0(X16,X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X16)))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~l1_orders_2(X16))&(r2_hidden(esk4_4(X16,X17,X18,X19),X17)|~r2_hidden(X19,X18)|~m1_subset_1(X19,u1_struct_0(X16))|X18!=k4_waybel_0(X16,X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X16)))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~l1_orders_2(X16))))&(~m1_subset_1(X21,u1_struct_0(X16))|(~r1_orders_2(X16,X21,X19)|~r2_hidden(X21,X17))|r2_hidden(X19,X18)|~m1_subset_1(X19,u1_struct_0(X16))|X18!=k4_waybel_0(X16,X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X16)))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~l1_orders_2(X16)))&((m1_subset_1(esk5_3(X16,X17,X18),u1_struct_0(X16))|X18=k4_waybel_0(X16,X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X16)))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~l1_orders_2(X16))&((~r2_hidden(esk5_3(X16,X17,X18),X18)|(~m1_subset_1(X23,u1_struct_0(X16))|(~r1_orders_2(X16,X23,esk5_3(X16,X17,X18))|~r2_hidden(X23,X17)))|X18=k4_waybel_0(X16,X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X16)))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~l1_orders_2(X16))&((m1_subset_1(esk6_3(X16,X17,X18),u1_struct_0(X16))|r2_hidden(esk5_3(X16,X17,X18),X18)|X18=k4_waybel_0(X16,X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X16)))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~l1_orders_2(X16))&((r1_orders_2(X16,esk6_3(X16,X17,X18),esk5_3(X16,X17,X18))|r2_hidden(esk5_3(X16,X17,X18),X18)|X18=k4_waybel_0(X16,X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X16)))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~l1_orders_2(X16))&(r2_hidden(esk6_3(X16,X17,X18),X17)|r2_hidden(esk5_3(X16,X17,X18),X18)|X18=k4_waybel_0(X16,X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X16)))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~l1_orders_2(X16))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d16_waybel_0])])])])])).
fof(c_0_25, plain, ![X34, X35, X36, X37]:((~r1_lattice3(X34,X37,X36)|r1_lattice3(X34,X37,X35)|~r1_orders_2(X34,X35,X36)|~m1_subset_1(X36,u1_struct_0(X34))|~m1_subset_1(X35,u1_struct_0(X34))|(~v4_orders_2(X34)|~l1_orders_2(X34)))&(~r2_lattice3(X34,X37,X35)|r2_lattice3(X34,X37,X36)|~r1_orders_2(X34,X35,X36)|~m1_subset_1(X36,u1_struct_0(X34))|~m1_subset_1(X35,u1_struct_0(X34))|(~v4_orders_2(X34)|~l1_orders_2(X34)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_yellow_0])])])])).
cnf(c_0_26, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,X1)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_14]), c_0_15])])).
cnf(c_0_27, plain, (r2_hidden(esk4_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k4_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_28, plain, ![X9, X10, X11]:((~r1_lattice3(X9,k1_tarski(X11),X10)|r1_orders_2(X9,X10,X11)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|~l1_orders_2(X9))&((~r1_orders_2(X9,X10,X11)|r1_lattice3(X9,k1_tarski(X11),X10)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|~l1_orders_2(X9))&((~r2_lattice3(X9,k1_tarski(X11),X10)|r1_orders_2(X9,X11,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|~l1_orders_2(X9))&(~r1_orders_2(X9,X11,X10)|r2_lattice3(X9,k1_tarski(X11),X10)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|~l1_orders_2(X9))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])).
cnf(c_0_29, plain, (r1_lattice3(X1,X2,X4)|~r1_lattice3(X1,X2,X3)|~r1_orders_2(X1,X4,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk4_4(X1,esk2_0,X2,X3))|X2!=k4_waybel_0(X1,esk2_0)|~r2_hidden(X3,X2)|~m1_subset_1(esk4_4(X1,esk2_0,X2,X3),u1_struct_0(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_31, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_32, plain, (r1_lattice3(X1,k1_tarski(X3),X2)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33, plain, (r1_orders_2(X1,esk4_4(X1,X2,X3,X4),X4)|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k4_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, plain, (m1_subset_1(esk4_4(X1,X2,X3,X4),u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k4_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_35, negated_conjecture, (r1_lattice3(esk1_0,X1,esk3_0)|X2!=k4_waybel_0(X3,esk2_0)|~r2_hidden(X4,X2)|~r1_lattice3(esk1_0,X1,esk4_4(X3,esk2_0,X2,X4))|~m1_subset_1(esk4_4(X3,esk2_0,X2,X4),u1_struct_0(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_orders_2(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_14]), c_0_15]), c_0_31])])).
cnf(c_0_36, plain, (r1_lattice3(X1,k1_tarski(X2),esk4_4(X1,X3,X4,X2))|X4!=k4_waybel_0(X1,X3)|~r2_hidden(X2,X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])).
cnf(c_0_37, negated_conjecture, (r1_lattice3(esk1_0,k1_tarski(X1),esk3_0)|X2!=k4_waybel_0(esk1_0,esk2_0)|~r2_hidden(X1,X2)|~m1_subset_1(esk4_4(esk1_0,esk2_0,X2,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_21]), c_0_15])])).
fof(c_0_38, plain, ![X30, X31]:(~l1_orders_2(X30)|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))|m1_subset_1(k4_waybel_0(X30,X31),k1_zfmisc_1(u1_struct_0(X30)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_waybel_0])])).
cnf(c_0_39, negated_conjecture, (r1_lattice3(esk1_0,k1_tarski(X1),esk3_0)|X2!=k4_waybel_0(esk1_0,esk2_0)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_34]), c_0_21]), c_0_15])])).
cnf(c_0_40, plain, (m1_subset_1(k4_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_41, negated_conjecture, (r1_lattice3(esk1_0,k1_tarski(X1),esk3_0)|k4_waybel_0(esk1_0,X2)!=k4_waybel_0(esk1_0,esk2_0)|~r2_hidden(X1,k4_waybel_0(esk1_0,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_15])])).
cnf(c_0_42, plain, (r2_hidden(esk7_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_43, plain, (r1_orders_2(X1,X3,X2)|~r1_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_44, negated_conjecture, (r1_lattice3(esk1_0,k1_tarski(esk7_3(X1,k4_waybel_0(esk1_0,X2),X3)),esk3_0)|r1_lattice3(X1,k4_waybel_0(esk1_0,X2),X3)|k4_waybel_0(esk1_0,X2)!=k4_waybel_0(esk1_0,esk2_0)|~m1_subset_1(esk7_3(X1,k4_waybel_0(esk1_0,X2),X3),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_45, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk7_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_46, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk7_3(X1,k4_waybel_0(esk1_0,X2),X3))|r1_lattice3(X1,k4_waybel_0(esk1_0,X2),X3)|k4_waybel_0(esk1_0,X2)!=k4_waybel_0(esk1_0,esk2_0)|~m1_subset_1(esk7_3(X1,k4_waybel_0(esk1_0,X2),X3),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_14]), c_0_15])])).
cnf(c_0_47, negated_conjecture, (r1_lattice3(esk1_0,k4_waybel_0(esk1_0,X1),esk3_0)|k4_waybel_0(esk1_0,X1)!=k4_waybel_0(esk1_0,esk2_0)|~m1_subset_1(esk7_3(esk1_0,k4_waybel_0(esk1_0,X1),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_14]), c_0_15])])).
cnf(c_0_48, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_49, negated_conjecture, (~r1_lattice3(esk1_0,esk2_0,esk3_0)|~r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_50, negated_conjecture, (r1_lattice3(esk1_0,k4_waybel_0(esk1_0,X1),esk3_0)|k4_waybel_0(esk1_0,X1)!=k4_waybel_0(esk1_0,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_14]), c_0_15])])).
cnf(c_0_51, negated_conjecture, (~r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_23])])).
cnf(c_0_52, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_50]), c_0_21])]), c_0_51]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 53
# Proof object clause steps            : 36
# Proof object formula steps           : 17
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 37
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 32
# Processed clauses                    : 229
# ...of these trivial                  : 1
# ...subsumed                          : 31
# ...remaining for further processing  : 197
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 15
# Backward-rewritten                   : 3
# Generated clauses                    : 371
# ...of the previous two non-trivial   : 359
# Contextual simplify-reflections      : 17
# Paramodulations                      : 368
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 147
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 139
# Current number of unprocessed clauses: 190
# ...number of literals in the above   : 2192
# Current number of archived formulas  : 0
# Current number of archived clauses   : 50
# Clause-clause subsumption calls (NU) : 12402
# Rec. Clause-clause subsumption calls : 283
# Non-unit clause-clause subsumptions  : 63
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 22903

# -------------------------------------------------
# User time                : 0.025 s
# System time              : 0.004 s
# Total time               : 0.029 s
# Maximum resident set size: 28624 pages
