# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/22ce8a886ae74990b2143d6207d5dd12.E_input', l54_zfmisc_1)).
fof(d1_relat_1, axiom, ![X1]:(v1_relat_1(X1)<=>![X2]:~((r2_hidden(X2,X1)&![X3, X4]:~(X2=k4_tarski(X3,X4))))), file('data/MPTP/proofs/22ce8a886ae74990b2143d6207d5dd12.E_input', d1_relat_1)).
fof(t10_mcart_1, conjecture, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(r2_hidden(k1_mcart_1(X1),X2)&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/22ce8a886ae74990b2143d6207d5dd12.E_input', t10_mcart_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/22ce8a886ae74990b2143d6207d5dd12.E_input', cc1_relset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/22ce8a886ae74990b2143d6207d5dd12.E_input', t3_subset)).
fof(l1_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/22ce8a886ae74990b2143d6207d5dd12.E_input', l1_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/22ce8a886ae74990b2143d6207d5dd12.E_input', t69_enumset1)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/22ce8a886ae74990b2143d6207d5dd12.E_input', t7_mcart_1)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/22ce8a886ae74990b2143d6207d5dd12.E_input', d2_tarski)).
fof(c_0_9, plain, ![X5, X6, X7, X8]:(((r2_hidden(X5,X7)|~r2_hidden(k4_tarski(X5,X6),k2_zfmisc_1(X7,X8)))&(r2_hidden(X6,X8)|~r2_hidden(k4_tarski(X5,X6),k2_zfmisc_1(X7,X8))))&(~r2_hidden(X5,X7)|~r2_hidden(X6,X8)|r2_hidden(k4_tarski(X5,X6),k2_zfmisc_1(X7,X8)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
fof(c_0_10, plain, ![X9, X10, X13, X15, X16]:((~v1_relat_1(X9)|(~r2_hidden(X10,X9)|X10=k4_tarski(esk1_2(X9,X10),esk2_2(X9,X10))))&((r2_hidden(esk3_1(X13),X13)|v1_relat_1(X13))&(esk3_1(X13)!=k4_tarski(X15,X16)|v1_relat_1(X13)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_relat_1])])])])])])])).
fof(c_0_11, negated_conjecture, ~(![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(r2_hidden(k1_mcart_1(X1),X2)&r2_hidden(k2_mcart_1(X1),X3)))), inference(assume_negation,[status(cth)],[t10_mcart_1])).
fof(c_0_12, plain, ![X20, X21, X22]:(~m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(X20,X21)))|v1_relat_1(X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_13, plain, ![X23, X24]:((~m1_subset_1(X23,k1_zfmisc_1(X24))|r1_tarski(X23,X24))&(~r1_tarski(X23,X24)|m1_subset_1(X23,k1_zfmisc_1(X24)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_14, plain, ![X25, X26]:((~r1_tarski(k1_tarski(X25),X26)|r2_hidden(X25,X26))&(~r2_hidden(X25,X26)|r1_tarski(k1_tarski(X25),X26))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])).
fof(c_0_15, plain, ![X27]:k2_tarski(X27,X27)=k1_tarski(X27), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_16, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, plain, (X2=k4_tarski(esk1_2(X1,X2),esk2_2(X1,X2))|~v1_relat_1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_18, negated_conjecture, (r2_hidden(esk4_0,k2_zfmisc_1(esk5_0,esk6_0))&(~r2_hidden(k1_mcart_1(esk4_0),esk5_0)|~r2_hidden(k2_mcart_1(esk4_0),esk6_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
fof(c_0_19, plain, ![X28, X29]:(k1_mcart_1(k4_tarski(X28,X29))=X28&k2_mcart_1(k4_tarski(X28,X29))=X29), inference(variable_rename,[status(thm)],[t7_mcart_1])).
cnf(c_0_20, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, plain, (r2_hidden(esk1_2(X1,X2),X3)|~v1_relat_1(X1)|~r2_hidden(X2,k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_25, negated_conjecture, (r2_hidden(esk4_0,k2_zfmisc_1(esk5_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (v1_relat_1(X1)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_28, plain, (r1_tarski(k2_tarski(X1,X1),X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_22, c_0_23])).
fof(c_0_29, plain, ![X30, X31, X32, X33, X34, X35, X36, X37]:(((~r2_hidden(X33,X32)|(X33=X30|X33=X31)|X32!=k2_tarski(X30,X31))&((X34!=X30|r2_hidden(X34,X32)|X32!=k2_tarski(X30,X31))&(X34!=X31|r2_hidden(X34,X32)|X32!=k2_tarski(X30,X31))))&(((esk7_3(X35,X36,X37)!=X35|~r2_hidden(esk7_3(X35,X36,X37),X37)|X37=k2_tarski(X35,X36))&(esk7_3(X35,X36,X37)!=X36|~r2_hidden(esk7_3(X35,X36,X37),X37)|X37=k2_tarski(X35,X36)))&(r2_hidden(esk7_3(X35,X36,X37),X37)|(esk7_3(X35,X36,X37)=X35|esk7_3(X35,X36,X37)=X36)|X37=k2_tarski(X35,X36)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_30, negated_conjecture, (r2_hidden(esk1_2(X1,esk4_0),esk5_0)|~v1_relat_1(X1)|~r2_hidden(esk4_0,X1)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31, plain, (esk1_2(X1,X2)=k1_mcart_1(X2)|~v1_relat_1(X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_26, c_0_17])).
cnf(c_0_32, plain, (v1_relat_1(k2_tarski(X1,X1))|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_33, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,X2))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_35, negated_conjecture, (r2_hidden(k1_mcart_1(esk4_0),esk5_0)|~v1_relat_1(X1)|~r2_hidden(esk4_0,X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_36, negated_conjecture, (v1_relat_1(k2_tarski(esk4_0,esk4_0))), inference(spm,[status(thm)],[c_0_32, c_0_25])).
cnf(c_0_37, plain, (r2_hidden(X1,k2_tarski(X1,X2))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_33])])).
cnf(c_0_38, plain, (r2_hidden(esk2_2(X1,X2),X3)|~v1_relat_1(X1)|~r2_hidden(X2,k2_zfmisc_1(X4,X3))|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_34, c_0_17])).
cnf(c_0_39, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_40, negated_conjecture, (~r2_hidden(k1_mcart_1(esk4_0),esk5_0)|~r2_hidden(k2_mcart_1(esk4_0),esk6_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_41, negated_conjecture, (r2_hidden(k1_mcart_1(esk4_0),esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37])])).
cnf(c_0_42, negated_conjecture, (r2_hidden(esk2_2(X1,esk4_0),esk6_0)|~v1_relat_1(X1)|~r2_hidden(esk4_0,X1)), inference(spm,[status(thm)],[c_0_38, c_0_25])).
cnf(c_0_43, plain, (esk2_2(X1,X2)=k2_mcart_1(X2)|~v1_relat_1(X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_39, c_0_17])).
cnf(c_0_44, negated_conjecture, (~r2_hidden(k2_mcart_1(esk4_0),esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_41])])).
cnf(c_0_45, negated_conjecture, (~v1_relat_1(X1)|~r2_hidden(esk4_0,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])).
cnf(c_0_46, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_36]), c_0_37])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 47
# Proof object clause steps            : 28
# Proof object formula steps           : 19
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 10
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 21
# Processed clauses                    : 63
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 63
# Other redundant clauses eliminated   : 11
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 2
# Generated clauses                    : 63
# ...of the previous two non-trivial   : 47
# Contextual simplify-reflections      : 0
# Paramodulations                      : 52
# Factorizations                       : 2
# Equation resolutions                 : 11
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 35
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 27
# Current number of unprocessed clauses: 26
# ...number of literals in the above   : 89
# Current number of archived formulas  : 0
# Current number of archived clauses   : 26
# Clause-clause subsumption calls (NU) : 50
# Rec. Clause-clause subsumption calls : 42
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2076

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.000 s
# Total time               : 0.009 s
# Maximum resident set size: 28648 pages
