# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/231ed609f5164f90bf4a6278f526c6a3.E_input', d1_tarski)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/231ed609f5164f90bf4a6278f526c6a3.E_input', t69_enumset1)).
fof(t143_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k9_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k1_relat_1(X3))&(r2_hidden(k4_tarski(X4,X1),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/231ed609f5164f90bf4a6278f526c6a3.E_input', t143_relat_1)).
fof(t117_funct_1, conjecture, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r2_hidden(X1,k1_relat_1(X2))=>k11_relat_1(X2,X1)=k1_tarski(k1_funct_1(X2,X1)))), file('data/MPTP/proofs/231ed609f5164f90bf4a6278f526c6a3.E_input', t117_funct_1)).
fof(d16_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:k11_relat_1(X1,X2)=k9_relat_1(X1,k1_tarski(X2))), file('data/MPTP/proofs/231ed609f5164f90bf4a6278f526c6a3.E_input', d16_relat_1)).
fof(d12_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2, X3]:(X3=k9_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(X5,k1_relat_1(X1))&(r2_hidden(X5,X2)&X4=k1_funct_1(X1,X5)))))), file('data/MPTP/proofs/231ed609f5164f90bf4a6278f526c6a3.E_input', d12_funct_1)).
fof(t8_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>(r2_hidden(X1,k1_relat_1(X3))&X2=k1_funct_1(X3,X1)))), file('data/MPTP/proofs/231ed609f5164f90bf4a6278f526c6a3.E_input', t8_funct_1)).
fof(c_0_7, plain, ![X6, X7, X8, X9, X10, X11]:(((~r2_hidden(X8,X7)|X8=X6|X7!=k1_tarski(X6))&(X9!=X6|r2_hidden(X9,X7)|X7!=k1_tarski(X6)))&((~r2_hidden(esk1_2(X10,X11),X11)|esk1_2(X10,X11)!=X10|X11=k1_tarski(X10))&(r2_hidden(esk1_2(X10,X11),X11)|esk1_2(X10,X11)=X10|X11=k1_tarski(X10)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_8, plain, ![X13]:k2_tarski(X13,X13)=k1_tarski(X13), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_9, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11, plain, (X1=X3|X2!=k2_tarski(X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_9, c_0_10])).
fof(c_0_12, plain, ![X14, X15, X16, X18]:(((r2_hidden(esk2_3(X14,X15,X16),k1_relat_1(X16))|~r2_hidden(X14,k9_relat_1(X16,X15))|~v1_relat_1(X16))&((r2_hidden(k4_tarski(esk2_3(X14,X15,X16),X14),X16)|~r2_hidden(X14,k9_relat_1(X16,X15))|~v1_relat_1(X16))&(r2_hidden(esk2_3(X14,X15,X16),X15)|~r2_hidden(X14,k9_relat_1(X16,X15))|~v1_relat_1(X16))))&(~r2_hidden(X18,k1_relat_1(X16))|(~r2_hidden(k4_tarski(X18,X14),X16)|~r2_hidden(X18,X15))|r2_hidden(X14,k9_relat_1(X16,X15))|~v1_relat_1(X16))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t143_relat_1])])])])])).
fof(c_0_13, negated_conjecture, ~(![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r2_hidden(X1,k1_relat_1(X2))=>k11_relat_1(X2,X1)=k1_tarski(k1_funct_1(X2,X1))))), inference(assume_negation,[status(cth)],[t117_funct_1])).
cnf(c_0_14, plain, (X1=X2|~r2_hidden(X1,k2_tarski(X2,X2))), inference(er,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (r2_hidden(esk2_3(X1,X2,X3),X2)|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_16, plain, ![X21, X22]:(~v1_relat_1(X21)|k11_relat_1(X21,X22)=k9_relat_1(X21,k1_tarski(X22))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d16_relat_1])])])).
fof(c_0_17, negated_conjecture, ((v1_relat_1(esk4_0)&v1_funct_1(esk4_0))&(r2_hidden(esk3_0,k1_relat_1(esk4_0))&k11_relat_1(esk4_0,esk3_0)!=k1_tarski(k1_funct_1(esk4_0,esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
cnf(c_0_18, plain, (r2_hidden(k4_tarski(esk2_3(X1,X2,X3),X1),X3)|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, plain, (esk2_3(X1,k2_tarski(X2,X2),X3)=X2|~v1_relat_1(X3)|~r2_hidden(X1,k9_relat_1(X3,k2_tarski(X2,X2)))), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_20, plain, (k11_relat_1(X1,X2)=k9_relat_1(X1,k1_tarski(X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, negated_conjecture, (k11_relat_1(esk4_0,esk3_0)!=k1_tarski(k1_funct_1(esk4_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, plain, (r2_hidden(esk1_2(X1,X2),X2)|esk1_2(X1,X2)=X1|X2=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_23, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~v1_relat_1(X3)|~r2_hidden(X2,k9_relat_1(X3,k2_tarski(X1,X1)))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_24, plain, (k11_relat_1(X1,X2)=k9_relat_1(X1,k2_tarski(X2,X2))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_20, c_0_10])).
cnf(c_0_25, negated_conjecture, (k11_relat_1(esk4_0,esk3_0)!=k2_tarski(k1_funct_1(esk4_0,esk3_0),k1_funct_1(esk4_0,esk3_0))), inference(rw,[status(thm)],[c_0_21, c_0_10])).
cnf(c_0_26, plain, (X2=k2_tarski(X1,X1)|esk1_2(X1,X2)=X1|r2_hidden(esk1_2(X1,X2),X2)), inference(rw,[status(thm)],[c_0_22, c_0_10])).
fof(c_0_27, plain, ![X23, X24, X25, X26, X28, X29, X30, X31, X33]:((((r2_hidden(esk5_4(X23,X24,X25,X26),k1_relat_1(X23))|~r2_hidden(X26,X25)|X25!=k9_relat_1(X23,X24)|(~v1_relat_1(X23)|~v1_funct_1(X23)))&((r2_hidden(esk5_4(X23,X24,X25,X26),X24)|~r2_hidden(X26,X25)|X25!=k9_relat_1(X23,X24)|(~v1_relat_1(X23)|~v1_funct_1(X23)))&(X26=k1_funct_1(X23,esk5_4(X23,X24,X25,X26))|~r2_hidden(X26,X25)|X25!=k9_relat_1(X23,X24)|(~v1_relat_1(X23)|~v1_funct_1(X23)))))&(~r2_hidden(X29,k1_relat_1(X23))|(~r2_hidden(X29,X24)|X28!=k1_funct_1(X23,X29))|r2_hidden(X28,X25)|X25!=k9_relat_1(X23,X24)|(~v1_relat_1(X23)|~v1_funct_1(X23))))&((~r2_hidden(esk6_3(X23,X30,X31),X31)|(~r2_hidden(X33,k1_relat_1(X23))|(~r2_hidden(X33,X30)|esk6_3(X23,X30,X31)!=k1_funct_1(X23,X33)))|X31=k9_relat_1(X23,X30)|(~v1_relat_1(X23)|~v1_funct_1(X23)))&((r2_hidden(esk7_3(X23,X30,X31),k1_relat_1(X23))|r2_hidden(esk6_3(X23,X30,X31),X31)|X31=k9_relat_1(X23,X30)|(~v1_relat_1(X23)|~v1_funct_1(X23)))&((r2_hidden(esk7_3(X23,X30,X31),X30)|r2_hidden(esk6_3(X23,X30,X31),X31)|X31=k9_relat_1(X23,X30)|(~v1_relat_1(X23)|~v1_funct_1(X23)))&(esk6_3(X23,X30,X31)=k1_funct_1(X23,esk7_3(X23,X30,X31))|r2_hidden(esk6_3(X23,X30,X31),X31)|X31=k9_relat_1(X23,X30)|(~v1_relat_1(X23)|~v1_funct_1(X23))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_funct_1])])])])])])).
fof(c_0_28, plain, ![X35, X36, X37]:(((r2_hidden(X35,k1_relat_1(X37))|~r2_hidden(k4_tarski(X35,X36),X37)|(~v1_relat_1(X37)|~v1_funct_1(X37)))&(X36=k1_funct_1(X37,X35)|~r2_hidden(k4_tarski(X35,X36),X37)|(~v1_relat_1(X37)|~v1_funct_1(X37))))&(~r2_hidden(X35,k1_relat_1(X37))|X36!=k1_funct_1(X37,X35)|r2_hidden(k4_tarski(X35,X36),X37)|(~v1_relat_1(X37)|~v1_funct_1(X37)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_funct_1])])])).
cnf(c_0_29, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~v1_relat_1(X3)|~r2_hidden(X2,k11_relat_1(X3,X1))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_30, negated_conjecture, (esk1_2(k1_funct_1(esk4_0,esk3_0),k11_relat_1(esk4_0,esk3_0))=k1_funct_1(esk4_0,esk3_0)|r2_hidden(esk1_2(k1_funct_1(esk4_0,esk3_0),k11_relat_1(esk4_0,esk3_0)),k11_relat_1(esk4_0,esk3_0))), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26])])).
cnf(c_0_31, negated_conjecture, (v1_relat_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_32, plain, (r2_hidden(X4,X5)|~r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(X1,X3)|X4!=k1_funct_1(X2,X1)|X5!=k9_relat_1(X2,X3)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_34, plain, (X2=k1_tarski(X1)|~r2_hidden(esk1_2(X1,X2),X2)|esk1_2(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_35, plain, (X1=k1_funct_1(X2,X3)|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36, negated_conjecture, (esk1_2(k1_funct_1(esk4_0,esk3_0),k11_relat_1(esk4_0,esk3_0))=k1_funct_1(esk4_0,esk3_0)|r2_hidden(k4_tarski(esk3_0,esk1_2(k1_funct_1(esk4_0,esk3_0),k11_relat_1(esk4_0,esk3_0))),esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])])).
cnf(c_0_37, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_38, plain, (r2_hidden(k1_funct_1(X1,X2),k9_relat_1(X1,X3))|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))|~r2_hidden(X2,X3)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_32])])).
cnf(c_0_39, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X2)), inference(rw,[status(thm)],[c_0_33, c_0_10])).
cnf(c_0_40, plain, (X2=k2_tarski(X1,X1)|esk1_2(X1,X2)!=X1|~r2_hidden(esk1_2(X1,X2),X2)), inference(rw,[status(thm)],[c_0_34, c_0_10])).
cnf(c_0_41, negated_conjecture, (esk1_2(k1_funct_1(esk4_0,esk3_0),k11_relat_1(esk4_0,esk3_0))=k1_funct_1(esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37]), c_0_31])])).
cnf(c_0_42, plain, (r2_hidden(k1_funct_1(X1,X2),k11_relat_1(X1,X3))|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X2,k2_tarski(X3,X3))|~r2_hidden(X2,k1_relat_1(X1))), inference(spm,[status(thm)],[c_0_38, c_0_24])).
cnf(c_0_43, plain, (r2_hidden(X1,k2_tarski(X1,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_39])])).
cnf(c_0_44, negated_conjecture, (~r2_hidden(k1_funct_1(esk4_0,esk3_0),k11_relat_1(esk4_0,esk3_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_25])).
cnf(c_0_45, plain, (r2_hidden(k1_funct_1(X1,X2),k11_relat_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk3_0,k1_relat_1(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_47, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_37]), c_0_31]), c_0_46])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 48
# Proof object clause steps            : 33
# Proof object formula steps           : 15
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 24
# Processed clauses                    : 80
# ...of these trivial                  : 0
# ...subsumed                          : 4
# ...remaining for further processing  : 76
# Other redundant clauses eliminated   : 13
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 223
# ...of the previous two non-trivial   : 206
# Contextual simplify-reflections      : 0
# Paramodulations                      : 212
# Factorizations                       : 0
# Equation resolutions                 : 13
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 43
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 36
# Current number of unprocessed clauses: 172
# ...number of literals in the above   : 844
# Current number of archived formulas  : 0
# Current number of archived clauses   : 27
# Clause-clause subsumption calls (NU) : 590
# Rec. Clause-clause subsumption calls : 185
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 8
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5990

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.004 s
# Total time               : 0.012 s
# Maximum resident set size: 28632 pages
