# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(fraenkel_a_3_0_waybel_9, axiom, ![X1, X2, X3, X4]:(((~(v2_struct_0(X2))&l1_struct_0(X2))&((~(v2_struct_0(X3))&l1_waybel_0(X3,X2))&m1_subset_1(X4,u1_struct_0(X3))))=>(r2_hidden(X1,a_3_0_waybel_9(X2,X3,X4))<=>?[X5]:(m1_subset_1(X5,u1_struct_0(X3))&(X1=X5&r1_orders_2(X3,X4,X5))))), file('data/MPTP/proofs/233841e160c6449d96b1b7ae7a764cfe.E_input', fraenkel_a_3_0_waybel_9)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/233841e160c6449d96b1b7ae7a764cfe.E_input', d3_tarski)).
fof(cc3_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/233841e160c6449d96b1b7ae7a764cfe.E_input', cc3_relset_1)).
fof(dt_u1_waybel_0, axiom, ![X1, X2]:((l1_struct_0(X1)&l1_waybel_0(X2,X1))=>(v1_funct_1(u1_waybel_0(X1,X2))&(v1_funct_2(u1_waybel_0(X1,X2),u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(u1_waybel_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))), file('data/MPTP/proofs/233841e160c6449d96b1b7ae7a764cfe.E_input', dt_u1_waybel_0)).
fof(fc15_yellow_6, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&(~(v2_struct_0(X2))&l1_waybel_0(X2,X1)))=>(v1_funct_1(u1_waybel_0(X1,X2))&(~(v1_xboole_0(u1_waybel_0(X1,X2)))&v1_funct_2(u1_waybel_0(X1,X2),u1_struct_0(X2),u1_struct_0(X1))))), file('data/MPTP/proofs/233841e160c6449d96b1b7ae7a764cfe.E_input', fc15_yellow_6)).
fof(t13_waybel_9, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>r1_tarski(u1_struct_0(k4_waybel_9(X1,X2,X3)),u1_struct_0(X2))))), file('data/MPTP/proofs/233841e160c6449d96b1b7ae7a764cfe.E_input', t13_waybel_9)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/233841e160c6449d96b1b7ae7a764cfe.E_input', t2_subset)).
fof(t12_waybel_9, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>u1_struct_0(k4_waybel_9(X1,X2,X3))=a_3_0_waybel_9(X1,X2,X3)))), file('data/MPTP/proofs/233841e160c6449d96b1b7ae7a764cfe.E_input', t12_waybel_9)).
fof(c_0_8, plain, ![X6, X7, X8, X9, X11]:(((m1_subset_1(esk1_4(X6,X7,X8,X9),u1_struct_0(X8))|~r2_hidden(X6,a_3_0_waybel_9(X7,X8,X9))|(v2_struct_0(X7)|~l1_struct_0(X7)|(v2_struct_0(X8)|~l1_waybel_0(X8,X7)|~m1_subset_1(X9,u1_struct_0(X8)))))&((X6=esk1_4(X6,X7,X8,X9)|~r2_hidden(X6,a_3_0_waybel_9(X7,X8,X9))|(v2_struct_0(X7)|~l1_struct_0(X7)|(v2_struct_0(X8)|~l1_waybel_0(X8,X7)|~m1_subset_1(X9,u1_struct_0(X8)))))&(r1_orders_2(X8,X9,esk1_4(X6,X7,X8,X9))|~r2_hidden(X6,a_3_0_waybel_9(X7,X8,X9))|(v2_struct_0(X7)|~l1_struct_0(X7)|(v2_struct_0(X8)|~l1_waybel_0(X8,X7)|~m1_subset_1(X9,u1_struct_0(X8)))))))&(~m1_subset_1(X11,u1_struct_0(X8))|(X6!=X11|~r1_orders_2(X8,X9,X11))|r2_hidden(X6,a_3_0_waybel_9(X7,X8,X9))|(v2_struct_0(X7)|~l1_struct_0(X7)|(v2_struct_0(X8)|~l1_waybel_0(X8,X7)|~m1_subset_1(X9,u1_struct_0(X8)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_3_0_waybel_9])])])])])])).
fof(c_0_9, plain, ![X12, X13, X14, X15, X16]:((~r1_tarski(X12,X13)|(~r2_hidden(X14,X12)|r2_hidden(X14,X13)))&((r2_hidden(esk2_2(X15,X16),X15)|r1_tarski(X15,X16))&(~r2_hidden(esk2_2(X15,X16),X16)|r1_tarski(X15,X16)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_10, plain, ![X18, X19, X20]:(~v1_xboole_0(X18)|(~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(X18,X19)))|v1_xboole_0(X20))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_relset_1])])])).
fof(c_0_11, plain, ![X21, X22]:((v1_funct_1(u1_waybel_0(X21,X22))|(~l1_struct_0(X21)|~l1_waybel_0(X22,X21)))&((v1_funct_2(u1_waybel_0(X21,X22),u1_struct_0(X22),u1_struct_0(X21))|(~l1_struct_0(X21)|~l1_waybel_0(X22,X21)))&(m1_subset_1(u1_waybel_0(X21,X22),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(X21))))|(~l1_struct_0(X21)|~l1_waybel_0(X22,X21))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_waybel_0])])])).
cnf(c_0_12, plain, (X1=esk1_4(X1,X2,X3,X4)|v2_struct_0(X2)|v2_struct_0(X3)|~r2_hidden(X1,a_3_0_waybel_9(X2,X3,X4))|~l1_struct_0(X2)|~l1_waybel_0(X3,X2)|~m1_subset_1(X4,u1_struct_0(X3))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_14, plain, ![X23, X24]:((v1_funct_1(u1_waybel_0(X23,X24))|(v2_struct_0(X23)|~l1_struct_0(X23)|(v2_struct_0(X24)|~l1_waybel_0(X24,X23))))&((~v1_xboole_0(u1_waybel_0(X23,X24))|(v2_struct_0(X23)|~l1_struct_0(X23)|(v2_struct_0(X24)|~l1_waybel_0(X24,X23))))&(v1_funct_2(u1_waybel_0(X23,X24),u1_struct_0(X24),u1_struct_0(X23))|(v2_struct_0(X23)|~l1_struct_0(X23)|(v2_struct_0(X24)|~l1_waybel_0(X24,X23)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc15_yellow_6])])])])).
cnf(c_0_15, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (m1_subset_1(u1_waybel_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_17, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>r1_tarski(u1_struct_0(k4_waybel_9(X1,X2,X3)),u1_struct_0(X2)))))), inference(assume_negation,[status(cth)],[t13_waybel_9])).
fof(c_0_18, plain, ![X28, X29]:(~m1_subset_1(X28,X29)|(v1_xboole_0(X29)|r2_hidden(X28,X29))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_19, plain, (m1_subset_1(esk1_4(X1,X2,X3,X4),u1_struct_0(X3))|v2_struct_0(X2)|v2_struct_0(X3)|~r2_hidden(X1,a_3_0_waybel_9(X2,X3,X4))|~l1_struct_0(X2)|~l1_waybel_0(X3,X2)|~m1_subset_1(X4,u1_struct_0(X3))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_20, plain, (esk1_4(esk2_2(a_3_0_waybel_9(X1,X2,X3),X4),X1,X2,X3)=esk2_2(a_3_0_waybel_9(X1,X2,X3),X4)|r1_tarski(a_3_0_waybel_9(X1,X2,X3),X4)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_21, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~v1_xboole_0(u1_waybel_0(X1,X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (v1_xboole_0(u1_waybel_0(X1,X2))|~v1_xboole_0(u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
fof(c_0_23, negated_conjecture, ((~v2_struct_0(esk3_0)&l1_struct_0(esk3_0))&((~v2_struct_0(esk4_0)&l1_waybel_0(esk4_0,esk3_0))&(m1_subset_1(esk5_0,u1_struct_0(esk4_0))&~r1_tarski(u1_struct_0(k4_waybel_9(esk3_0,esk4_0,esk5_0)),u1_struct_0(esk4_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])).
fof(c_0_24, plain, ![X30, X31, X32]:(v2_struct_0(X30)|~l1_struct_0(X30)|(v2_struct_0(X31)|~l1_waybel_0(X31,X30)|(~m1_subset_1(X32,u1_struct_0(X31))|u1_struct_0(k4_waybel_9(X30,X31,X32))=a_3_0_waybel_9(X30,X31,X32)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t12_waybel_9])])])])).
cnf(c_0_25, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, plain, (r1_tarski(a_3_0_waybel_9(X1,X2,X3),X4)|m1_subset_1(esk2_2(a_3_0_waybel_9(X1,X2,X3),X4),u1_struct_0(X2))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_13])).
cnf(c_0_27, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~v1_xboole_0(u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28, negated_conjecture, (~r1_tarski(u1_struct_0(k4_waybel_9(esk3_0,esk4_0,esk5_0)),u1_struct_0(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (v2_struct_0(X1)|v2_struct_0(X2)|u1_struct_0(k4_waybel_9(X1,X2,X3))=a_3_0_waybel_9(X1,X2,X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31, negated_conjecture, (l1_waybel_0(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, negated_conjecture, (l1_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_33, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_34, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_35, plain, (r1_tarski(X1,X2)|~r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_36, plain, (r1_tarski(a_3_0_waybel_9(X1,X2,X3),X4)|r2_hidden(esk2_2(a_3_0_waybel_9(X1,X2,X3),X4),u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])).
cnf(c_0_37, negated_conjecture, (~r1_tarski(a_3_0_waybel_9(esk3_0,esk4_0,esk5_0),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30]), c_0_31]), c_0_32])]), c_0_33]), c_0_34])).
cnf(c_0_38, plain, (r1_tarski(a_3_0_waybel_9(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_39, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_30]), c_0_31]), c_0_32])]), c_0_34]), c_0_33]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 40
# Proof object clause steps            : 23
# Proof object formula steps           : 17
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 14
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 22
# Processed clauses                    : 71
# ...of these trivial                  : 0
# ...subsumed                          : 5
# ...remaining for further processing  : 66
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 2
# Generated clauses                    : 50
# ...of the previous two non-trivial   : 47
# Contextual simplify-reflections      : 4
# Paramodulations                      : 49
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 42
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 33
# Current number of unprocessed clauses: 17
# ...number of literals in the above   : 191
# Current number of archived formulas  : 0
# Current number of archived clauses   : 23
# Clause-clause subsumption calls (NU) : 895
# Rec. Clause-clause subsumption calls : 36
# Non-unit clause-clause subsumptions  : 10
# Unit Clause-clause subsumption calls : 5
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3723

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.006 s
# Total time               : 0.011 s
# Maximum resident set size: 28788 pages
