# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t51_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((r1_yellow_0(X1,X2)|r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))=>k1_yellow_0(X1,X2)=k1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))), file('data/MPTP/proofs/236a87d1ae0c48e5ac8cd13d227e3f31.E_input', t51_yellow_0)).
fof(t46_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2, X3]:((![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)<=>r2_lattice3(X1,X3,X4)))&r1_yellow_0(X1,X2))=>r1_yellow_0(X1,X3))), file('data/MPTP/proofs/236a87d1ae0c48e5ac8cd13d227e3f31.E_input', t46_yellow_0)).
fof(t5_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r2_lattice3(X1,X2,X3)=>r2_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3))&((r2_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3)=>r2_lattice3(X1,X2,X3))&((r1_lattice3(X1,X2,X3)=>r1_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3))&(r1_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3)=>r1_lattice3(X1,X2,X3))))))), file('data/MPTP/proofs/236a87d1ae0c48e5ac8cd13d227e3f31.E_input', t5_yellow_0)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/236a87d1ae0c48e5ac8cd13d227e3f31.E_input', dt_k1_yellow_0)).
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/236a87d1ae0c48e5ac8cd13d227e3f31.E_input', d9_yellow_0)).
fof(c_0_5, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((r1_yellow_0(X1,X2)|r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))=>k1_yellow_0(X1,X2)=k1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))))), inference(assume_negation,[status(cth)],[t51_yellow_0])).
fof(c_0_6, plain, ![X7, X8, X9]:((m1_subset_1(esk3_3(X7,X8,X9),u1_struct_0(X7))|~r1_yellow_0(X7,X8)|r1_yellow_0(X7,X9)|(v2_struct_0(X7)|~l1_orders_2(X7)))&((~r2_lattice3(X7,X8,esk3_3(X7,X8,X9))|~r2_lattice3(X7,X9,esk3_3(X7,X8,X9))|~r1_yellow_0(X7,X8)|r1_yellow_0(X7,X9)|(v2_struct_0(X7)|~l1_orders_2(X7)))&(r2_lattice3(X7,X8,esk3_3(X7,X8,X9))|r2_lattice3(X7,X9,esk3_3(X7,X8,X9))|~r1_yellow_0(X7,X8)|r1_yellow_0(X7,X9)|(v2_struct_0(X7)|~l1_orders_2(X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t46_yellow_0])])])])])])).
fof(c_0_7, negated_conjecture, ((~v2_struct_0(esk1_0)&l1_orders_2(esk1_0))&((r1_yellow_0(esk1_0,esk2_0)|r1_yellow_0(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0))))&k1_yellow_0(esk1_0,esk2_0)!=k1_yellow_0(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_5])])])])).
fof(c_0_8, plain, ![X11, X12, X13]:((~r2_lattice3(X11,X12,X13)|r2_lattice3(X11,k3_xboole_0(X12,u1_struct_0(X11)),X13)|~m1_subset_1(X13,u1_struct_0(X11))|(v2_struct_0(X11)|~l1_orders_2(X11)))&((~r2_lattice3(X11,k3_xboole_0(X12,u1_struct_0(X11)),X13)|r2_lattice3(X11,X12,X13)|~m1_subset_1(X13,u1_struct_0(X11))|(v2_struct_0(X11)|~l1_orders_2(X11)))&((~r1_lattice3(X11,X12,X13)|r1_lattice3(X11,k3_xboole_0(X12,u1_struct_0(X11)),X13)|~m1_subset_1(X13,u1_struct_0(X11))|(v2_struct_0(X11)|~l1_orders_2(X11)))&(~r1_lattice3(X11,k3_xboole_0(X12,u1_struct_0(X11)),X13)|r1_lattice3(X11,X12,X13)|~m1_subset_1(X13,u1_struct_0(X11))|(v2_struct_0(X11)|~l1_orders_2(X11)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t5_yellow_0])])])])])).
cnf(c_0_9, plain, (m1_subset_1(esk3_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X3)|v2_struct_0(X1)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, negated_conjecture, (r1_yellow_0(esk1_0,esk2_0)|r1_yellow_0(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, plain, (r2_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r2_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk3_3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),X1),u1_struct_0(esk1_0))|r1_yellow_0(esk1_0,esk2_0)|r1_yellow_0(esk1_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11])]), c_0_12])).
cnf(c_0_15, plain, (r2_lattice3(X1,X2,esk3_3(X1,X2,X3))|r2_lattice3(X1,X3,esk3_3(X1,X2,X3))|r1_yellow_0(X1,X3)|v2_struct_0(X1)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_16, negated_conjecture, (r2_lattice3(esk1_0,X1,esk3_3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),X2))|r1_yellow_0(esk1_0,esk2_0)|r1_yellow_0(esk1_0,X2)|~r2_lattice3(esk1_0,k3_xboole_0(X1,u1_struct_0(esk1_0)),esk3_3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),X2))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_11])]), c_0_12])).
cnf(c_0_17, negated_conjecture, (r2_lattice3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),esk3_3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),X1))|r2_lattice3(esk1_0,X1,esk3_3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),X1))|r1_yellow_0(esk1_0,esk2_0)|r1_yellow_0(esk1_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_10]), c_0_11])]), c_0_12])).
cnf(c_0_18, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk3_3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),X1))|r2_lattice3(esk1_0,X1,esk3_3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),X1))|r1_yellow_0(esk1_0,esk2_0)|r1_yellow_0(esk1_0,X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_19, plain, (r2_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3)|v2_struct_0(X1)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_20, plain, (r1_yellow_0(X1,X3)|v2_struct_0(X1)|~r2_lattice3(X1,X2,esk3_3(X1,X2,X3))|~r2_lattice3(X1,X3,esk3_3(X1,X2,X3))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_21, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk3_3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),esk2_0))|r1_yellow_0(esk1_0,esk2_0)), inference(ef,[status(thm)],[c_0_18])).
fof(c_0_22, plain, ![X14, X15]:(~l1_orders_2(X14)|m1_subset_1(k1_yellow_0(X14,X15),u1_struct_0(X14))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
fof(c_0_23, plain, ![X16, X17, X18, X19]:(((r2_lattice3(X16,X17,X18)|X18!=k1_yellow_0(X16,X17)|~r1_yellow_0(X16,X17)|~m1_subset_1(X18,u1_struct_0(X16))|~l1_orders_2(X16))&(~m1_subset_1(X19,u1_struct_0(X16))|(~r2_lattice3(X16,X17,X19)|r1_orders_2(X16,X18,X19))|X18!=k1_yellow_0(X16,X17)|~r1_yellow_0(X16,X17)|~m1_subset_1(X18,u1_struct_0(X16))|~l1_orders_2(X16)))&((m1_subset_1(esk4_3(X16,X17,X18),u1_struct_0(X16))|~r2_lattice3(X16,X17,X18)|X18=k1_yellow_0(X16,X17)|~r1_yellow_0(X16,X17)|~m1_subset_1(X18,u1_struct_0(X16))|~l1_orders_2(X16))&((r2_lattice3(X16,X17,esk4_3(X16,X17,X18))|~r2_lattice3(X16,X17,X18)|X18=k1_yellow_0(X16,X17)|~r1_yellow_0(X16,X17)|~m1_subset_1(X18,u1_struct_0(X16))|~l1_orders_2(X16))&(~r1_orders_2(X16,X18,esk4_3(X16,X17,X18))|~r2_lattice3(X16,X17,X18)|X18=k1_yellow_0(X16,X17)|~r1_yellow_0(X16,X17)|~m1_subset_1(X18,u1_struct_0(X16))|~l1_orders_2(X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
cnf(c_0_24, negated_conjecture, (r2_lattice3(esk1_0,k3_xboole_0(X1,u1_struct_0(esk1_0)),esk3_3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),X2))|r1_yellow_0(esk1_0,esk2_0)|r1_yellow_0(esk1_0,X2)|~r2_lattice3(esk1_0,X1,esk3_3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),X2))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_14]), c_0_11])]), c_0_12])).
cnf(c_0_25, negated_conjecture, (r1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),esk3_3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),esk2_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_11])]), c_0_12]), c_0_10])).
cnf(c_0_26, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27, plain, (r2_lattice3(X1,X2,X3)|X3!=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28, negated_conjecture, (r1_yellow_0(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_21]), c_0_25])).
cnf(c_0_29, plain, (r2_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),k1_yellow_0(X1,X3))|v2_struct_0(X1)|~r2_lattice3(X1,X2,k1_yellow_0(X1,X3))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_19, c_0_26])).
cnf(c_0_30, plain, (r2_lattice3(X1,X2,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_27]), c_0_26])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk3_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|r1_yellow_0(esk1_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_28]), c_0_11])]), c_0_12])).
cnf(c_0_32, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X1))|X3=k1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_33, plain, (r2_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),k1_yellow_0(X1,X2))|v2_struct_0(X1)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_34, negated_conjecture, (r2_lattice3(esk1_0,k3_xboole_0(X1,u1_struct_0(esk1_0)),esk3_3(esk1_0,esk2_0,X2))|r1_yellow_0(esk1_0,X2)|~r2_lattice3(esk1_0,X1,esk3_3(esk1_0,esk2_0,X2))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_31]), c_0_11])]), c_0_12])).
cnf(c_0_35, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk3_3(esk1_0,esk2_0,X1))|r2_lattice3(esk1_0,X1,esk3_3(esk1_0,esk2_0,X1))|r1_yellow_0(esk1_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_28]), c_0_11])]), c_0_12])).
cnf(c_0_36, plain, (k1_yellow_0(X1,X2)=k1_yellow_0(X1,X3)|m1_subset_1(esk4_3(X1,X3,k1_yellow_0(X1,X2)),u1_struct_0(X1))|~r2_lattice3(X1,X3,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X3)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_32, c_0_26])).
cnf(c_0_37, negated_conjecture, (r2_lattice3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),k1_yellow_0(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_28]), c_0_11])]), c_0_12])).
cnf(c_0_38, negated_conjecture, (k1_yellow_0(esk1_0,esk2_0)!=k1_yellow_0(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_39, plain, (r2_lattice3(X1,X2,esk4_3(X1,X2,X3))|X3=k1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_40, negated_conjecture, (r2_lattice3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),esk3_3(esk1_0,esk2_0,X1))|r2_lattice3(esk1_0,X1,esk3_3(esk1_0,esk2_0,X1))|r1_yellow_0(esk1_0,X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_41, negated_conjecture, (m1_subset_1(esk4_3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),k1_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_11])]), c_0_38])).
cnf(c_0_42, plain, (k1_yellow_0(X1,X2)=k1_yellow_0(X1,X3)|r2_lattice3(X1,X3,esk4_3(X1,X3,k1_yellow_0(X1,X2)))|~r2_lattice3(X1,X3,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X3)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_39, c_0_26])).
cnf(c_0_43, negated_conjecture, (r2_lattice3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),esk3_3(esk1_0,esk2_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0))))|r1_yellow_0(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)))), inference(ef,[status(thm)],[c_0_40])).
cnf(c_0_44, plain, (r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~r1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_45, negated_conjecture, (r2_lattice3(esk1_0,X1,esk4_3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),k1_yellow_0(esk1_0,esk2_0)))|~r2_lattice3(esk1_0,k3_xboole_0(X1,u1_struct_0(esk1_0)),esk4_3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),k1_yellow_0(esk1_0,esk2_0)))|~r1_yellow_0(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_41]), c_0_11])]), c_0_12])).
cnf(c_0_46, negated_conjecture, (r2_lattice3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),esk4_3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),k1_yellow_0(esk1_0,esk2_0)))|~r1_yellow_0(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_37]), c_0_11])]), c_0_38])).
cnf(c_0_47, negated_conjecture, (r2_lattice3(esk1_0,X1,esk3_3(esk1_0,esk2_0,X2))|r1_yellow_0(esk1_0,X2)|~r2_lattice3(esk1_0,k3_xboole_0(X1,u1_struct_0(esk1_0)),esk3_3(esk1_0,esk2_0,X2))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_31]), c_0_11])]), c_0_12])).
cnf(c_0_48, negated_conjecture, (r1_yellow_0(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)))|~r2_lattice3(esk1_0,esk2_0,esk3_3(esk1_0,esk2_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_43]), c_0_28]), c_0_11])]), c_0_12])).
cnf(c_0_49, plain, (X2=k1_yellow_0(X1,X3)|~r1_orders_2(X1,X2,esk4_3(X1,X3,X2))|~r2_lattice3(X1,X3,X2)|~r1_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_50, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_44]), c_0_26])).
cnf(c_0_51, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),k1_yellow_0(esk1_0,esk2_0)))|~r1_yellow_0(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_52, negated_conjecture, (r1_yellow_0(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_43]), c_0_48])).
cnf(c_0_53, plain, (k1_yellow_0(X1,X2)=k1_yellow_0(X1,X3)|~r2_lattice3(X1,X2,esk4_3(X1,X3,k1_yellow_0(X1,X2)))|~r2_lattice3(X1,X3,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X3)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_26]), c_0_36])).
cnf(c_0_54, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)),k1_yellow_0(esk1_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_52])])).
cnf(c_0_55, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_37]), c_0_52]), c_0_28]), c_0_11])]), c_0_38]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 56
# Proof object clause steps            : 45
# Proof object formula steps           : 11
# Proof object conjectures             : 30
# Proof object clause conjectures      : 27
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 29
# Proof object simplifying inferences  : 58
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 149
# ...of these trivial                  : 2
# ...subsumed                          : 16
# ...remaining for further processing  : 131
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 26
# Generated clauses                    : 198
# ...of the previous two non-trivial   : 176
# Contextual simplify-reflections      : 7
# Paramodulations                      : 188
# Factorizations                       : 8
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 88
#    Positive orientable unit clauses  : 29
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 57
# Current number of unprocessed clauses: 45
# ...number of literals in the above   : 141
# Current number of archived formulas  : 0
# Current number of archived clauses   : 43
# Clause-clause subsumption calls (NU) : 800
# Rec. Clause-clause subsumption calls : 254
# Non-unit clause-clause subsumptions  : 23
# Unit Clause-clause subsumption calls : 92
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 179
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 10740

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.004 s
# Total time               : 0.015 s
# Maximum resident set size: 28620 pages
