# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_wellord2, axiom, ![X1, X2]:(v1_relat_1(X2)=>(X2=k1_wellord2(X1)<=>(k3_relat_1(X2)=X1&![X3, X4]:((r2_hidden(X3,X1)&r2_hidden(X4,X1))=>(r2_hidden(k4_tarski(X3,X4),X2)<=>r1_tarski(X3,X4)))))), file('data/MPTP/proofs/23f867688ba04153a6295f7fb78adf88.E_input', d1_wellord2)).
fof(dt_k1_wellord2, axiom, ![X1]:v1_relat_1(k1_wellord2(X1)), file('data/MPTP/proofs/23f867688ba04153a6295f7fb78adf88.E_input', dt_k1_wellord2)).
fof(t13_wellord1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k1_wellord1(X2,X1),k3_relat_1(X2))), file('data/MPTP/proofs/23f867688ba04153a6295f7fb78adf88.E_input', t13_wellord1)).
fof(t11_wellord2, conjecture, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r4_wellord1(k1_wellord2(X1),k1_wellord2(X2))=>X1=X2))), file('data/MPTP/proofs/23f867688ba04153a6295f7fb78adf88.E_input', t11_wellord2)).
fof(t57_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v2_wellord1(X1)=>![X2]:~((r2_hidden(X2,k3_relat_1(X1))&r4_wellord1(X1,k2_wellord1(X1,k1_wellord1(X1,X2))))))), file('data/MPTP/proofs/23f867688ba04153a6295f7fb78adf88.E_input', t57_wellord1)).
fof(t8_wellord2, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_wellord1(k1_wellord2(X2),X1)=k1_wellord2(X1)), file('data/MPTP/proofs/23f867688ba04153a6295f7fb78adf88.E_input', t8_wellord2)).
fof(t26_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r1_ordinal1(X1,X2)|r2_hidden(X2,X1)))), file('data/MPTP/proofs/23f867688ba04153a6295f7fb78adf88.E_input', t26_ordinal1)).
fof(t10_wellord2, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)=>X1=k1_wellord1(k1_wellord2(X2),X1)))), file('data/MPTP/proofs/23f867688ba04153a6295f7fb78adf88.E_input', t10_wellord2)).
fof(t7_wellord2, axiom, ![X1]:(v3_ordinal1(X1)=>v2_wellord1(k1_wellord2(X1))), file('data/MPTP/proofs/23f867688ba04153a6295f7fb78adf88.E_input', t7_wellord2)).
fof(t50_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r4_wellord1(X1,X2)=>r4_wellord1(X2,X1)))), file('data/MPTP/proofs/23f867688ba04153a6295f7fb78adf88.E_input', t50_wellord1)).
fof(redefinition_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)<=>r1_tarski(X1,X2))), file('data/MPTP/proofs/23f867688ba04153a6295f7fb78adf88.E_input', redefinition_r1_ordinal1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/23f867688ba04153a6295f7fb78adf88.E_input', d10_xboole_0)).
fof(c_0_12, plain, ![X16, X17, X18, X19]:(((k3_relat_1(X17)=X16|X17!=k1_wellord2(X16)|~v1_relat_1(X17))&((~r2_hidden(k4_tarski(X18,X19),X17)|r1_tarski(X18,X19)|(~r2_hidden(X18,X16)|~r2_hidden(X19,X16))|X17!=k1_wellord2(X16)|~v1_relat_1(X17))&(~r1_tarski(X18,X19)|r2_hidden(k4_tarski(X18,X19),X17)|(~r2_hidden(X18,X16)|~r2_hidden(X19,X16))|X17!=k1_wellord2(X16)|~v1_relat_1(X17))))&(((r2_hidden(esk3_2(X16,X17),X16)|k3_relat_1(X17)!=X16|X17=k1_wellord2(X16)|~v1_relat_1(X17))&(r2_hidden(esk4_2(X16,X17),X16)|k3_relat_1(X17)!=X16|X17=k1_wellord2(X16)|~v1_relat_1(X17)))&((~r2_hidden(k4_tarski(esk3_2(X16,X17),esk4_2(X16,X17)),X17)|~r1_tarski(esk3_2(X16,X17),esk4_2(X16,X17))|k3_relat_1(X17)!=X16|X17=k1_wellord2(X16)|~v1_relat_1(X17))&(r2_hidden(k4_tarski(esk3_2(X16,X17),esk4_2(X16,X17)),X17)|r1_tarski(esk3_2(X16,X17),esk4_2(X16,X17))|k3_relat_1(X17)!=X16|X17=k1_wellord2(X16)|~v1_relat_1(X17))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_wellord2])])])])])).
fof(c_0_13, plain, ![X15]:v1_relat_1(k1_wellord2(X15)), inference(variable_rename,[status(thm)],[dt_k1_wellord2])).
fof(c_0_14, plain, ![X9, X10]:(~v1_relat_1(X10)|r1_tarski(k1_wellord1(X10,X9),k3_relat_1(X10))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_wellord1])])).
cnf(c_0_15, plain, (k3_relat_1(X1)=X2|X1!=k1_wellord2(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, plain, (v1_relat_1(k1_wellord2(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_17, negated_conjecture, ~(![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r4_wellord1(k1_wellord2(X1),k1_wellord2(X2))=>X1=X2)))), inference(assume_negation,[status(cth)],[t11_wellord2])).
fof(c_0_18, plain, ![X13, X14]:(~v1_relat_1(X13)|(~v2_wellord1(X13)|(~r2_hidden(X14,k3_relat_1(X13))|~r4_wellord1(X13,k2_wellord1(X13,k1_wellord1(X13,X14)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t57_wellord1])])])).
fof(c_0_19, plain, ![X28, X29]:(~r1_tarski(X28,X29)|k2_wellord1(k1_wellord2(X29),X28)=k1_wellord2(X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_wellord2])])).
cnf(c_0_20, plain, (r1_tarski(k1_wellord1(X1,X2),k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (k3_relat_1(k1_wellord2(X1))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_15]), c_0_16])])).
fof(c_0_22, plain, ![X24, X25]:(~v3_ordinal1(X24)|(~v3_ordinal1(X25)|(r1_ordinal1(X24,X25)|r2_hidden(X25,X24)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t26_ordinal1])])])).
fof(c_0_23, negated_conjecture, (v3_ordinal1(esk1_0)&(v3_ordinal1(esk2_0)&(r4_wellord1(k1_wellord2(esk1_0),k1_wellord2(esk2_0))&esk1_0!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).
fof(c_0_24, plain, ![X26, X27]:(~v3_ordinal1(X26)|(~v3_ordinal1(X27)|(~r2_hidden(X26,X27)|X26=k1_wellord1(k1_wellord2(X27),X26)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_wellord2])])])).
fof(c_0_25, plain, ![X30]:(~v3_ordinal1(X30)|v2_wellord1(k1_wellord2(X30))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_wellord2])])).
cnf(c_0_26, plain, (~v1_relat_1(X1)|~v2_wellord1(X1)|~r2_hidden(X2,k3_relat_1(X1))|~r4_wellord1(X1,k2_wellord1(X1,k1_wellord1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, plain, (k2_wellord1(k1_wellord2(X2),X1)=k1_wellord2(X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, plain, (r1_tarski(k1_wellord1(k1_wellord2(X1),X2),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_16])])).
fof(c_0_29, plain, ![X22, X23]:(~v1_relat_1(X22)|(~v1_relat_1(X23)|(~r4_wellord1(X22,X23)|r4_wellord1(X23,X22)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t50_wellord1])])])).
cnf(c_0_30, plain, (r1_ordinal1(X1,X2)|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, negated_conjecture, (v3_ordinal1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, plain, (X1=k1_wellord1(k1_wellord2(X2),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, plain, (v2_wellord1(k1_wellord2(X1))|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, plain, (~r2_hidden(X1,X2)|~v2_wellord1(k1_wellord2(X2))|~r4_wellord1(k1_wellord2(X2),k1_wellord2(k1_wellord1(k1_wellord2(X2),X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_16])]), c_0_21]), c_0_28])])).
cnf(c_0_35, plain, (r4_wellord1(X2,X1)|~v1_relat_1(X1)|~v1_relat_1(X2)|~r4_wellord1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, negated_conjecture, (r4_wellord1(k1_wellord2(esk1_0),k1_wellord2(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_37, plain, ![X11, X12]:((~r1_ordinal1(X11,X12)|r1_tarski(X11,X12)|(~v3_ordinal1(X11)|~v3_ordinal1(X12)))&(~r1_tarski(X11,X12)|r1_ordinal1(X11,X12)|(~v3_ordinal1(X11)|~v3_ordinal1(X12)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r1_ordinal1])])])).
cnf(c_0_38, negated_conjecture, (r2_hidden(esk2_0,X1)|r1_ordinal1(X1,esk2_0)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_39, negated_conjecture, (v3_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_40, plain, (~r2_hidden(X1,X2)|~r4_wellord1(k1_wellord2(X2),k2_wellord1(k1_wellord2(X2),X1))|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_32]), c_0_16])]), c_0_21])]), c_0_33])).
cnf(c_0_41, plain, (~r2_hidden(X1,X2)|~r4_wellord1(k1_wellord2(X2),k1_wellord2(X1))|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_32]), c_0_33])).
cnf(c_0_42, negated_conjecture, (r4_wellord1(k1_wellord2(esk2_0),k1_wellord2(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_16]), c_0_16])])).
fof(c_0_43, plain, ![X7, X8]:(((r1_tarski(X7,X8)|X7!=X8)&(r1_tarski(X8,X7)|X7!=X8))&(~r1_tarski(X7,X8)|~r1_tarski(X8,X7)|X7=X8)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_44, plain, (r1_tarski(X1,X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_45, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|r1_ordinal1(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_46, plain, (~r2_hidden(X1,X2)|~r1_tarski(X1,X2)|~r4_wellord1(k1_wellord2(X2),k1_wellord2(X1))|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_40, c_0_27])).
cnf(c_0_47, negated_conjecture, (r2_hidden(esk1_0,X1)|r1_ordinal1(X1,esk1_0)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_30, c_0_39])).
cnf(c_0_48, negated_conjecture, (~r2_hidden(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_31]), c_0_39])])).
cnf(c_0_49, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_50, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|r1_tarski(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_31]), c_0_39])])).
cnf(c_0_51, negated_conjecture, (esk1_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_52, negated_conjecture, (~r2_hidden(esk2_0,esk1_0)|~r1_tarski(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_36]), c_0_39]), c_0_31])])).
cnf(c_0_53, negated_conjecture, (r1_ordinal1(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_31]), c_0_48])).
cnf(c_0_54, negated_conjecture, (~r1_tarski(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51]), c_0_52])).
cnf(c_0_55, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_53]), c_0_39]), c_0_31])]), c_0_54]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 56
# Proof object clause steps            : 31
# Proof object formula steps           : 25
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 34
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 24
# Processed clauses                    : 69
# ...of these trivial                  : 1
# ...subsumed                          : 0
# ...remaining for further processing  : 68
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 44
# ...of the previous two non-trivial   : 35
# Contextual simplify-reflections      : 3
# Paramodulations                      : 35
# Factorizations                       : 0
# Equation resolutions                 : 9
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 42
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 30
# Current number of unprocessed clauses: 13
# ...number of literals in the above   : 38
# Current number of archived formulas  : 0
# Current number of archived clauses   : 24
# Clause-clause subsumption calls (NU) : 214
# Rec. Clause-clause subsumption calls : 68
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 61
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2413

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28640 pages
