# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t34_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(X2=k16_lattice3(X1,X3)<=>(r3_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_lattice3(X1,X4,X3)=>r3_lattices(X1,X4,X2))))))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', t34_lattice3)).
fof(dt_k16_lattice3, axiom, ![X1, X2]:((~(v2_struct_0(X1))&l3_lattices(X1))=>m1_subset_1(k16_lattice3(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', dt_k16_lattice3)).
fof(t37_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:k15_lattice3(X1,X2)=k16_lattice3(X1,a_2_2_lattice3(X1,X2))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', t37_lattice3)).
fof(dt_k15_lattice3, axiom, ![X1, X2]:((~(v2_struct_0(X1))&l3_lattices(X1))=>m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', dt_k15_lattice3)).
fof(t43_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k15_lattice3(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k16_lattice3(X1,k6_domain_1(u1_struct_0(X1),X2))=X2))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', t43_lattice3)).
fof(t50_lattice3, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))&k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', t50_lattice3)).
fof(t48_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X2)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>~((r3_lattices(X1,X4,X5)&r2_hidden(X5,X3)))))))=>r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3)))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', t48_lattice3)).
fof(t47_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(k15_lattice3(X1,X2)=k15_lattice3(X1,a_2_5_lattice3(X1,X2))&k16_lattice3(X1,X2)=k16_lattice3(X1,a_2_6_lattice3(X1,X2)))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', t47_lattice3)).
fof(fraenkel_a_2_6_lattice3, axiom, ![X1, X2, X3]:((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))=>(r2_hidden(X1,a_2_6_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r3_lattices(X2,X5,X4)&r2_hidden(X5,X3))))))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', fraenkel_a_2_6_lattice3)).
fof(t46_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(r1_tarski(X2,X3)=>(r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3))&r3_lattices(X1,k16_lattice3(X1,X3),k16_lattice3(X1,X2))))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', t46_lattice3)).
fof(d16_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r3_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X2,X4)))))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', d16_lattice3)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', t2_xboole_1)).
fof(fraenkel_a_2_2_lattice3, axiom, ![X1, X2, X3]:((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))=>(r2_hidden(X1,a_2_2_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&r4_lattice3(X2,X4,X3))))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', fraenkel_a_2_2_lattice3)).
fof(d21_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X3=k15_lattice3(X1,X2)<=>(r4_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r4_lattice3(X1,X4,X2)=>r1_lattices(X1,X3,X4)))))))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', d21_lattice3)).
fof(d16_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v13_lattices(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(X2=k5_lattices(X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_lattices(X1,X2,X3)=X2&k2_lattices(X1,X3,X2)=X2)))))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', d16_lattices)).
fof(dt_k5_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>m1_subset_1(k5_lattices(X1),u1_struct_0(X1))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', dt_k5_lattices)).
fof(t21_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattices(X1,X2,X3)<=>k2_lattices(X1,X2,X3)=X2)))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', t21_lattices)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', cc1_lattices)).
fof(d6_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v6_lattices(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2))))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', d6_lattices)).
fof(dt_l3_lattices, axiom, ![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', dt_l3_lattices)).
fof(d13_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v13_lattices(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_lattices(X1,X2,X3)=X2&k2_lattices(X1,X3,X2)=X2))))), file('data/MPTP/proofs/241478103fe34aab87cb8d6434d60329.E_input', d13_lattices)).
fof(c_0_21, plain, ![X6, X7, X8, X9, X10]:(((r3_lattice3(X6,X7,X8)|X7!=k16_lattice3(X6,X8)|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6)))))&(~m1_subset_1(X9,u1_struct_0(X6))|(~r3_lattice3(X6,X9,X8)|r3_lattices(X6,X9,X7))|X7!=k16_lattice3(X6,X8)|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6))))))&((m1_subset_1(esk1_3(X6,X7,X10),u1_struct_0(X6))|~r3_lattice3(X6,X7,X10)|X7=k16_lattice3(X6,X10)|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6)))))&((r3_lattice3(X6,esk1_3(X6,X7,X10),X10)|~r3_lattice3(X6,X7,X10)|X7=k16_lattice3(X6,X10)|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6)))))&(~r3_lattices(X6,esk1_3(X6,X7,X10),X7)|~r3_lattice3(X6,X7,X10)|X7=k16_lattice3(X6,X10)|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_lattice3])])])])])])])).
fof(c_0_22, plain, ![X12, X13]:(v2_struct_0(X12)|~l3_lattices(X12)|m1_subset_1(k16_lattice3(X12,X13),u1_struct_0(X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k16_lattice3])])])).
cnf(c_0_23, plain, (r3_lattices(X2,X1,X4)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r3_lattice3(X2,X1,X3)|X4!=k16_lattice3(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_24, plain, (v2_struct_0(X1)|m1_subset_1(k16_lattice3(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_25, plain, ![X14, X15]:(v2_struct_0(X14)|(~v10_lattices(X14)|(~v4_lattice3(X14)|~l3_lattices(X14)))|k15_lattice3(X14,X15)=k16_lattice3(X14,a_2_2_lattice3(X14,X15))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t37_lattice3])])])])).
cnf(c_0_26, plain, (r3_lattice3(X1,X2,X3)|v2_struct_0(X1)|X2!=k16_lattice3(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (r3_lattices(X1,X2,k16_lattice3(X1,X3))|v2_struct_0(X1)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_23]), c_0_24])).
cnf(c_0_28, plain, (v2_struct_0(X1)|k15_lattice3(X1,X2)=k16_lattice3(X1,a_2_2_lattice3(X1,X2))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_29, plain, (r3_lattice3(X1,k16_lattice3(X1,X2),X2)|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_26]), c_0_24])).
fof(c_0_30, plain, ![X16, X17]:(v2_struct_0(X16)|~l3_lattices(X16)|m1_subset_1(k15_lattice3(X16,X17),u1_struct_0(X16))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k15_lattice3])])])).
cnf(c_0_31, plain, (r3_lattices(X1,X2,k15_lattice3(X1,X3))|v2_struct_0(X1)|~r3_lattice3(X1,X2,a_2_2_lattice3(X1,X3))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_32, plain, (r3_lattice3(X1,k15_lattice3(X1,X2),a_2_2_lattice3(X1,X2))|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_29, c_0_28])).
cnf(c_0_33, plain, (v2_struct_0(X1)|m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_34, plain, ![X18, X19]:((k15_lattice3(X18,k6_domain_1(u1_struct_0(X18),X19))=X19|~m1_subset_1(X19,u1_struct_0(X18))|(v2_struct_0(X18)|(~v10_lattices(X18)|(~v4_lattice3(X18)|~l3_lattices(X18)))))&(k16_lattice3(X18,k6_domain_1(u1_struct_0(X18),X19))=X19|~m1_subset_1(X19,u1_struct_0(X18))|(v2_struct_0(X18)|(~v10_lattices(X18)|(~v4_lattice3(X18)|~l3_lattices(X18)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t43_lattice3])])])])])).
fof(c_0_35, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))&k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0)))), inference(assume_negation,[status(cth)],[t50_lattice3])).
fof(c_0_36, plain, ![X21, X22, X23, X25]:((m1_subset_1(esk3_3(X21,X22,X23),u1_struct_0(X21))|r3_lattices(X21,k15_lattice3(X21,X22),k15_lattice3(X21,X23))|(v2_struct_0(X21)|(~v10_lattices(X21)|(~v4_lattice3(X21)|~l3_lattices(X21)))))&((r2_hidden(esk3_3(X21,X22,X23),X22)|r3_lattices(X21,k15_lattice3(X21,X22),k15_lattice3(X21,X23))|(v2_struct_0(X21)|(~v10_lattices(X21)|(~v4_lattice3(X21)|~l3_lattices(X21)))))&(~m1_subset_1(X25,u1_struct_0(X21))|(~r3_lattices(X21,esk3_3(X21,X22,X23),X25)|~r2_hidden(X25,X23))|r3_lattices(X21,k15_lattice3(X21,X22),k15_lattice3(X21,X23))|(v2_struct_0(X21)|(~v10_lattices(X21)|(~v4_lattice3(X21)|~l3_lattices(X21))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t48_lattice3])])])])])])).
cnf(c_0_37, plain, (r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X2))|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])).
cnf(c_0_38, plain, (k15_lattice3(X1,k6_domain_1(u1_struct_0(X1),X2))=X2|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_39, negated_conjecture, ((~v2_struct_0(esk2_0)&(v10_lattices(esk2_0)&(v4_lattice3(esk2_0)&l3_lattices(esk2_0))))&(v2_struct_0(esk2_0)|(~v10_lattices(esk2_0)|(~v13_lattices(esk2_0)|~l3_lattices(esk2_0)))|k5_lattices(esk2_0)!=k15_lattice3(esk2_0,k1_xboole_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_35])])])])).
fof(c_0_40, plain, ![X26, X27]:((k15_lattice3(X26,X27)=k15_lattice3(X26,a_2_5_lattice3(X26,X27))|(v2_struct_0(X26)|(~v10_lattices(X26)|(~v4_lattice3(X26)|~l3_lattices(X26)))))&(k16_lattice3(X26,X27)=k16_lattice3(X26,a_2_6_lattice3(X26,X27))|(v2_struct_0(X26)|(~v10_lattices(X26)|(~v4_lattice3(X26)|~l3_lattices(X26)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t47_lattice3])])])])])).
cnf(c_0_41, plain, (r3_lattices(X2,k15_lattice3(X2,X3),k15_lattice3(X2,X4))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r3_lattices(X2,esk3_3(X2,X3,X4),X1)|~r2_hidden(X1,X4)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_42, plain, (r3_lattices(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_43, plain, (m1_subset_1(esk3_3(X1,X2,X3),u1_struct_0(X1))|r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3))|v2_struct_0(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44, plain, (r2_hidden(esk3_3(X1,X2,X3),X2)|r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3))|v2_struct_0(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_45, negated_conjecture, (l3_lattices(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_46, negated_conjecture, (v4_lattice3(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_47, negated_conjecture, (v10_lattices(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_48, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_49, plain, (k16_lattice3(X1,X2)=k16_lattice3(X1,a_2_6_lattice3(X1,X2))|v2_struct_0(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_50, plain, ![X28, X29, X30, X33, X34, X35]:(((m1_subset_1(esk4_3(X28,X29,X30),u1_struct_0(X29))|~r2_hidden(X28,a_2_6_lattice3(X29,X30))|(v2_struct_0(X29)|(~v10_lattices(X29)|(~v4_lattice3(X29)|~l3_lattices(X29)))))&((X28=esk4_3(X28,X29,X30)|~r2_hidden(X28,a_2_6_lattice3(X29,X30))|(v2_struct_0(X29)|(~v10_lattices(X29)|(~v4_lattice3(X29)|~l3_lattices(X29)))))&((m1_subset_1(esk5_3(X28,X29,X30),u1_struct_0(X29))|~r2_hidden(X28,a_2_6_lattice3(X29,X30))|(v2_struct_0(X29)|(~v10_lattices(X29)|(~v4_lattice3(X29)|~l3_lattices(X29)))))&((r3_lattices(X29,esk5_3(X28,X29,X30),esk4_3(X28,X29,X30))|~r2_hidden(X28,a_2_6_lattice3(X29,X30))|(v2_struct_0(X29)|(~v10_lattices(X29)|(~v4_lattice3(X29)|~l3_lattices(X29)))))&(r2_hidden(esk5_3(X28,X29,X30),X30)|~r2_hidden(X28,a_2_6_lattice3(X29,X30))|(v2_struct_0(X29)|(~v10_lattices(X29)|(~v4_lattice3(X29)|~l3_lattices(X29)))))))))&(~m1_subset_1(X34,u1_struct_0(X29))|(X28!=X34|(~m1_subset_1(X35,u1_struct_0(X29))|(~r3_lattices(X29,X35,X34)|~r2_hidden(X35,X33))))|r2_hidden(X28,a_2_6_lattice3(X29,X33))|(v2_struct_0(X29)|(~v10_lattices(X29)|(~v4_lattice3(X29)|~l3_lattices(X29)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_6_lattice3])])])])])])])).
cnf(c_0_51, plain, (r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3))|v2_struct_0(X1)|~r2_hidden(esk3_3(X1,X2,X3),X3)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk3_3(esk2_0,X1,X2),X1)|r3_lattices(esk2_0,k15_lattice3(esk2_0,X1),k15_lattice3(esk2_0,X2))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46]), c_0_47])]), c_0_48])).
fof(c_0_53, plain, ![X36, X37, X38]:((r3_lattices(X36,k15_lattice3(X36,X37),k15_lattice3(X36,X38))|~r1_tarski(X37,X38)|(v2_struct_0(X36)|(~v10_lattices(X36)|(~v4_lattice3(X36)|~l3_lattices(X36)))))&(r3_lattices(X36,k16_lattice3(X36,X38),k16_lattice3(X36,X37))|~r1_tarski(X37,X38)|(v2_struct_0(X36)|(~v10_lattices(X36)|(~v4_lattice3(X36)|~l3_lattices(X36)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t46_lattice3])])])])])).
fof(c_0_54, plain, ![X39, X40, X41, X42, X43]:((~r3_lattice3(X39,X40,X41)|(~m1_subset_1(X42,u1_struct_0(X39))|(~r2_hidden(X42,X41)|r1_lattices(X39,X40,X42)))|~m1_subset_1(X40,u1_struct_0(X39))|(v2_struct_0(X39)|~l3_lattices(X39)))&((m1_subset_1(esk6_3(X39,X40,X43),u1_struct_0(X39))|r3_lattice3(X39,X40,X43)|~m1_subset_1(X40,u1_struct_0(X39))|(v2_struct_0(X39)|~l3_lattices(X39)))&((r2_hidden(esk6_3(X39,X40,X43),X43)|r3_lattice3(X39,X40,X43)|~m1_subset_1(X40,u1_struct_0(X39))|(v2_struct_0(X39)|~l3_lattices(X39)))&(~r1_lattices(X39,X40,esk6_3(X39,X40,X43))|r3_lattice3(X39,X40,X43)|~m1_subset_1(X40,u1_struct_0(X39))|(v2_struct_0(X39)|~l3_lattices(X39)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattice3])])])])])])])).
cnf(c_0_55, plain, (r3_lattice3(X1,k16_lattice3(X1,X2),a_2_6_lattice3(X1,X2))|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_29, c_0_49])).
cnf(c_0_56, plain, (r2_hidden(X3,a_2_6_lattice3(X2,X5))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~m1_subset_1(X4,u1_struct_0(X2))|~r3_lattices(X2,X4,X1)|~r2_hidden(X4,X5)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_57, negated_conjecture, (r3_lattices(esk2_0,k15_lattice3(esk2_0,X1),k15_lattice3(esk2_0,X1))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_45]), c_0_46]), c_0_47])]), c_0_48])).
cnf(c_0_58, plain, (r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3))|v2_struct_0(X1)|~r1_tarski(X2,X3)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_59, plain, (k15_lattice3(X1,X2)=k15_lattice3(X1,a_2_5_lattice3(X1,X2))|v2_struct_0(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_60, plain, ![X45]:r1_tarski(k1_xboole_0,X45), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_61, plain, (r1_lattices(X1,X2,X4)|v2_struct_0(X1)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_62, plain, (r3_lattice3(X1,k16_lattice3(X1,X2),a_2_6_lattice3(X1,a_2_6_lattice3(X1,X2)))|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_55, c_0_49])).
cnf(c_0_63, plain, (r2_hidden(X1,a_2_6_lattice3(X2,X3))|v2_struct_0(X2)|~r2_hidden(X4,X3)|~r3_lattices(X2,X4,X1)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)), inference(er,[status(thm)],[c_0_56])).
cnf(c_0_64, negated_conjecture, (r3_lattices(esk2_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_38]), c_0_45]), c_0_46]), c_0_47])]), c_0_48])).
cnf(c_0_65, plain, (r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3))|v2_struct_0(X1)|~r1_tarski(X2,a_2_5_lattice3(X1,X3))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_66, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_67, plain, (r1_lattices(X1,k16_lattice3(X1,X2),X3)|v2_struct_0(X1)|~r2_hidden(X3,a_2_6_lattice3(X1,a_2_6_lattice3(X1,X2)))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_24])).
cnf(c_0_68, negated_conjecture, (r2_hidden(X1,a_2_6_lattice3(esk2_0,X2))|~r2_hidden(X1,X2)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_45]), c_0_46]), c_0_47])]), c_0_48])).
cnf(c_0_69, plain, (r3_lattices(X1,k15_lattice3(X1,k1_xboole_0),k15_lattice3(X1,X2))|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_70, negated_conjecture, (r1_lattices(esk2_0,k16_lattice3(esk2_0,X1),X2)|~r2_hidden(X2,a_2_6_lattice3(esk2_0,X1))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_45]), c_0_46]), c_0_47])]), c_0_48])).
cnf(c_0_71, plain, (r3_lattices(X1,k15_lattice3(X1,k1_xboole_0),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_69, c_0_38])).
fof(c_0_72, plain, ![X46, X47, X48, X50, X51]:(((m1_subset_1(esk7_3(X46,X47,X48),u1_struct_0(X47))|~r2_hidden(X46,a_2_2_lattice3(X47,X48))|(v2_struct_0(X47)|(~v10_lattices(X47)|(~v4_lattice3(X47)|~l3_lattices(X47)))))&((X46=esk7_3(X46,X47,X48)|~r2_hidden(X46,a_2_2_lattice3(X47,X48))|(v2_struct_0(X47)|(~v10_lattices(X47)|(~v4_lattice3(X47)|~l3_lattices(X47)))))&(r4_lattice3(X47,esk7_3(X46,X47,X48),X48)|~r2_hidden(X46,a_2_2_lattice3(X47,X48))|(v2_struct_0(X47)|(~v10_lattices(X47)|(~v4_lattice3(X47)|~l3_lattices(X47)))))))&(~m1_subset_1(X51,u1_struct_0(X47))|(X46!=X51|~r4_lattice3(X47,X51,X50))|r2_hidden(X46,a_2_2_lattice3(X47,X50))|(v2_struct_0(X47)|(~v10_lattices(X47)|(~v4_lattice3(X47)|~l3_lattices(X47)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_2_lattice3])])])])])])])).
fof(c_0_73, plain, ![X52, X53, X54, X55]:(((r4_lattice3(X52,X54,X53)|X54!=k15_lattice3(X52,X53)|~m1_subset_1(X54,u1_struct_0(X52))|(v2_struct_0(X52)|(~v10_lattices(X52)|(~v4_lattice3(X52)|~l3_lattices(X52))))|(v2_struct_0(X52)|~l3_lattices(X52)))&(~m1_subset_1(X55,u1_struct_0(X52))|(~r4_lattice3(X52,X55,X53)|r1_lattices(X52,X54,X55))|X54!=k15_lattice3(X52,X53)|~m1_subset_1(X54,u1_struct_0(X52))|(v2_struct_0(X52)|(~v10_lattices(X52)|(~v4_lattice3(X52)|~l3_lattices(X52))))|(v2_struct_0(X52)|~l3_lattices(X52))))&((m1_subset_1(esk8_3(X52,X53,X54),u1_struct_0(X52))|~r4_lattice3(X52,X54,X53)|X54=k15_lattice3(X52,X53)|~m1_subset_1(X54,u1_struct_0(X52))|(v2_struct_0(X52)|(~v10_lattices(X52)|(~v4_lattice3(X52)|~l3_lattices(X52))))|(v2_struct_0(X52)|~l3_lattices(X52)))&((r4_lattice3(X52,esk8_3(X52,X53,X54),X53)|~r4_lattice3(X52,X54,X53)|X54=k15_lattice3(X52,X53)|~m1_subset_1(X54,u1_struct_0(X52))|(v2_struct_0(X52)|(~v10_lattices(X52)|(~v4_lattice3(X52)|~l3_lattices(X52))))|(v2_struct_0(X52)|~l3_lattices(X52)))&(~r1_lattices(X52,X54,esk8_3(X52,X53,X54))|~r4_lattice3(X52,X54,X53)|X54=k15_lattice3(X52,X53)|~m1_subset_1(X54,u1_struct_0(X52))|(v2_struct_0(X52)|(~v10_lattices(X52)|(~v4_lattice3(X52)|~l3_lattices(X52))))|(v2_struct_0(X52)|~l3_lattices(X52)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d21_lattice3])])])])])])).
cnf(c_0_74, negated_conjecture, (r1_lattices(esk2_0,k15_lattice3(esk2_0,X1),X2)|~r2_hidden(X2,a_2_6_lattice3(esk2_0,a_2_2_lattice3(esk2_0,X1)))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_28]), c_0_45]), c_0_46]), c_0_47])]), c_0_48])).
cnf(c_0_75, plain, (r2_hidden(X1,a_2_6_lattice3(X2,X3))|v2_struct_0(X2)|~r2_hidden(k15_lattice3(X2,k1_xboole_0),X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_71]), c_0_33])).
cnf(c_0_76, plain, (r2_hidden(X3,a_2_2_lattice3(X2,X4))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~r4_lattice3(X2,X1,X4)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_77, plain, (r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X1)|X2!=k15_lattice3(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_78, negated_conjecture, (r1_lattices(esk2_0,k15_lattice3(esk2_0,X1),X2)|~r2_hidden(k15_lattice3(esk2_0,k1_xboole_0),a_2_2_lattice3(esk2_0,X1))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_45]), c_0_46]), c_0_47])]), c_0_48])).
cnf(c_0_79, plain, (r2_hidden(X1,a_2_2_lattice3(X2,X3))|v2_struct_0(X2)|~r4_lattice3(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)), inference(er,[status(thm)],[c_0_76])).
cnf(c_0_80, plain, (v2_struct_0(X1)|r4_lattice3(X1,X2,X3)|X2!=k15_lattice3(X1,X3)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_77])).
cnf(c_0_81, negated_conjecture, (r1_lattices(esk2_0,k15_lattice3(esk2_0,X1),X2)|~r4_lattice3(esk2_0,k15_lattice3(esk2_0,k1_xboole_0),X1)|~m1_subset_1(k15_lattice3(esk2_0,k1_xboole_0),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_45]), c_0_46]), c_0_47])]), c_0_48])).
cnf(c_0_82, plain, (r4_lattice3(X1,k15_lattice3(X1,X2),X2)|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_80]), c_0_33])).
fof(c_0_83, plain, ![X57, X58, X59]:(((k2_lattices(X57,X58,X59)=X58|~m1_subset_1(X59,u1_struct_0(X57))|X58!=k5_lattices(X57)|~m1_subset_1(X58,u1_struct_0(X57))|~v13_lattices(X57)|(v2_struct_0(X57)|~l1_lattices(X57)))&(k2_lattices(X57,X59,X58)=X58|~m1_subset_1(X59,u1_struct_0(X57))|X58!=k5_lattices(X57)|~m1_subset_1(X58,u1_struct_0(X57))|~v13_lattices(X57)|(v2_struct_0(X57)|~l1_lattices(X57))))&((m1_subset_1(esk9_2(X57,X58),u1_struct_0(X57))|X58=k5_lattices(X57)|~m1_subset_1(X58,u1_struct_0(X57))|~v13_lattices(X57)|(v2_struct_0(X57)|~l1_lattices(X57)))&(k2_lattices(X57,X58,esk9_2(X57,X58))!=X58|k2_lattices(X57,esk9_2(X57,X58),X58)!=X58|X58=k5_lattices(X57)|~m1_subset_1(X58,u1_struct_0(X57))|~v13_lattices(X57)|(v2_struct_0(X57)|~l1_lattices(X57))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattices])])])])])])).
fof(c_0_84, plain, ![X61]:(v2_struct_0(X61)|~l1_lattices(X61)|m1_subset_1(k5_lattices(X61),u1_struct_0(X61))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_lattices])])])).
fof(c_0_85, plain, ![X62, X63, X64]:((~r1_lattices(X62,X63,X64)|k2_lattices(X62,X63,X64)=X63|~m1_subset_1(X64,u1_struct_0(X62))|~m1_subset_1(X63,u1_struct_0(X62))|(v2_struct_0(X62)|(~v8_lattices(X62)|(~v9_lattices(X62)|~l3_lattices(X62)))))&(k2_lattices(X62,X63,X64)!=X63|r1_lattices(X62,X63,X64)|~m1_subset_1(X64,u1_struct_0(X62))|~m1_subset_1(X63,u1_struct_0(X62))|(v2_struct_0(X62)|(~v8_lattices(X62)|(~v9_lattices(X62)|~l3_lattices(X62)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t21_lattices])])])])])).
cnf(c_0_86, negated_conjecture, (r1_lattices(esk2_0,k15_lattice3(esk2_0,X1),X2)|~r4_lattice3(esk2_0,k15_lattice3(esk2_0,k1_xboole_0),a_2_5_lattice3(esk2_0,X1))|~m1_subset_1(k15_lattice3(esk2_0,k1_xboole_0),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_59]), c_0_45]), c_0_46]), c_0_47])]), c_0_48])).
cnf(c_0_87, plain, (r4_lattice3(X1,k15_lattice3(X1,X2),a_2_5_lattice3(X1,X2))|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_82, c_0_59])).
cnf(c_0_88, plain, (k2_lattices(X1,X2,X3)=X3|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|X3!=k5_lattices(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v13_lattices(X1)|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_89, plain, (v2_struct_0(X1)|m1_subset_1(k5_lattices(X1),u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_90, plain, (k2_lattices(X1,X2,X3)=X2|v2_struct_0(X1)|~r1_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_91, negated_conjecture, (r1_lattices(esk2_0,k15_lattice3(esk2_0,k1_xboole_0),X1)|~m1_subset_1(k15_lattice3(esk2_0,k1_xboole_0),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_45]), c_0_46]), c_0_47])]), c_0_48])).
cnf(c_0_92, negated_conjecture, (v2_struct_0(esk2_0)|~v10_lattices(esk2_0)|~v13_lattices(esk2_0)|~l3_lattices(esk2_0)|k5_lattices(esk2_0)!=k15_lattice3(esk2_0,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_93, plain, (k2_lattices(X1,X2,k5_lattices(X1))=k5_lattices(X1)|v2_struct_0(X1)|~l1_lattices(X1)|~v13_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_88]), c_0_89])).
cnf(c_0_94, negated_conjecture, (k2_lattices(esk2_0,k15_lattice3(esk2_0,k1_xboole_0),X1)=k15_lattice3(esk2_0,k1_xboole_0)|~v9_lattices(esk2_0)|~v8_lattices(esk2_0)|~m1_subset_1(k15_lattice3(esk2_0,k1_xboole_0),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_45])]), c_0_48])).
cnf(c_0_95, negated_conjecture, (k5_lattices(esk2_0)!=k15_lattice3(esk2_0,k1_xboole_0)|~v13_lattices(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_47]), c_0_45])]), c_0_48])).
cnf(c_0_96, negated_conjecture, (~v9_lattices(esk2_0)|~v8_lattices(esk2_0)|~l1_lattices(esk2_0)|~v13_lattices(esk2_0)|~m1_subset_1(k15_lattice3(esk2_0,k1_xboole_0),u1_struct_0(esk2_0))|~m1_subset_1(k5_lattices(esk2_0),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_48]), c_0_95])).
fof(c_0_97, plain, ![X65]:((~v2_struct_0(X65)|(v2_struct_0(X65)|~v10_lattices(X65))|~l3_lattices(X65))&((v4_lattices(X65)|(v2_struct_0(X65)|~v10_lattices(X65))|~l3_lattices(X65))&((v5_lattices(X65)|(v2_struct_0(X65)|~v10_lattices(X65))|~l3_lattices(X65))&((v6_lattices(X65)|(v2_struct_0(X65)|~v10_lattices(X65))|~l3_lattices(X65))&((v7_lattices(X65)|(v2_struct_0(X65)|~v10_lattices(X65))|~l3_lattices(X65))&((v8_lattices(X65)|(v2_struct_0(X65)|~v10_lattices(X65))|~l3_lattices(X65))&(v9_lattices(X65)|(v2_struct_0(X65)|~v10_lattices(X65))|~l3_lattices(X65)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
fof(c_0_98, plain, ![X66, X67, X68]:((~v6_lattices(X66)|(~m1_subset_1(X67,u1_struct_0(X66))|(~m1_subset_1(X68,u1_struct_0(X66))|k2_lattices(X66,X67,X68)=k2_lattices(X66,X68,X67)))|(v2_struct_0(X66)|~l1_lattices(X66)))&((m1_subset_1(esk10_1(X66),u1_struct_0(X66))|v6_lattices(X66)|(v2_struct_0(X66)|~l1_lattices(X66)))&((m1_subset_1(esk11_1(X66),u1_struct_0(X66))|v6_lattices(X66)|(v2_struct_0(X66)|~l1_lattices(X66)))&(k2_lattices(X66,esk10_1(X66),esk11_1(X66))!=k2_lattices(X66,esk11_1(X66),esk10_1(X66))|v6_lattices(X66)|(v2_struct_0(X66)|~l1_lattices(X66)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_lattices])])])])])])).
fof(c_0_99, plain, ![X71]:((l1_lattices(X71)|~l3_lattices(X71))&(l2_lattices(X71)|~l3_lattices(X71))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
cnf(c_0_100, negated_conjecture, (~v9_lattices(esk2_0)|~v8_lattices(esk2_0)|~l1_lattices(esk2_0)|~v13_lattices(esk2_0)|~m1_subset_1(k15_lattice3(esk2_0,k1_xboole_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_89]), c_0_48])).
cnf(c_0_101, plain, (v9_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_102, plain, (k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2)|v2_struct_0(X1)|~v6_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_103, plain, (v6_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_104, plain, (l1_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_105, negated_conjecture, (~v8_lattices(esk2_0)|~l1_lattices(esk2_0)|~v13_lattices(esk2_0)|~m1_subset_1(k15_lattice3(esk2_0,k1_xboole_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_45]), c_0_47])]), c_0_48])).
cnf(c_0_106, plain, (v8_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_97])).
fof(c_0_107, plain, ![X72, X74, X75]:(((m1_subset_1(esk12_1(X72),u1_struct_0(X72))|~v13_lattices(X72)|(v2_struct_0(X72)|~l1_lattices(X72)))&((k2_lattices(X72,esk12_1(X72),X74)=esk12_1(X72)|~m1_subset_1(X74,u1_struct_0(X72))|~v13_lattices(X72)|(v2_struct_0(X72)|~l1_lattices(X72)))&(k2_lattices(X72,X74,esk12_1(X72))=esk12_1(X72)|~m1_subset_1(X74,u1_struct_0(X72))|~v13_lattices(X72)|(v2_struct_0(X72)|~l1_lattices(X72)))))&((m1_subset_1(esk13_2(X72,X75),u1_struct_0(X72))|~m1_subset_1(X75,u1_struct_0(X72))|v13_lattices(X72)|(v2_struct_0(X72)|~l1_lattices(X72)))&(k2_lattices(X72,X75,esk13_2(X72,X75))!=X75|k2_lattices(X72,esk13_2(X72,X75),X75)!=X75|~m1_subset_1(X75,u1_struct_0(X72))|v13_lattices(X72)|(v2_struct_0(X72)|~l1_lattices(X72))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_lattices])])])])])])).
cnf(c_0_108, plain, (k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_104])).
cnf(c_0_109, negated_conjecture, (~l1_lattices(esk2_0)|~v13_lattices(esk2_0)|~m1_subset_1(k15_lattice3(esk2_0,k1_xboole_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_45]), c_0_47])]), c_0_48])).
cnf(c_0_110, plain, (v13_lattices(X1)|v2_struct_0(X1)|k2_lattices(X1,X2,esk13_2(X1,X2))!=X2|k2_lattices(X1,esk13_2(X1,X2),X2)!=X2|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_111, negated_conjecture, (k2_lattices(esk2_0,X1,k15_lattice3(esk2_0,k1_xboole_0))=k15_lattice3(esk2_0,k1_xboole_0)|~v9_lattices(esk2_0)|~v8_lattices(esk2_0)|~m1_subset_1(k15_lattice3(esk2_0,k1_xboole_0),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_94]), c_0_45]), c_0_47])]), c_0_48])).
cnf(c_0_112, negated_conjecture, (~l1_lattices(esk2_0)|~v13_lattices(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_33]), c_0_45])]), c_0_48])).
cnf(c_0_113, negated_conjecture, (~v9_lattices(esk2_0)|~v8_lattices(esk2_0)|~l1_lattices(esk2_0)|~m1_subset_1(esk13_2(esk2_0,k15_lattice3(esk2_0,k1_xboole_0)),u1_struct_0(esk2_0))|~m1_subset_1(k15_lattice3(esk2_0,k1_xboole_0),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_48]), c_0_94]), c_0_112])).
cnf(c_0_114, plain, (m1_subset_1(esk13_2(X1,X2),u1_struct_0(X1))|v13_lattices(X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_115, negated_conjecture, (~v9_lattices(esk2_0)|~v8_lattices(esk2_0)|~l1_lattices(esk2_0)|~m1_subset_1(k15_lattice3(esk2_0,k1_xboole_0),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_48]), c_0_112])).
cnf(c_0_116, negated_conjecture, (~v8_lattices(esk2_0)|~l1_lattices(esk2_0)|~m1_subset_1(k15_lattice3(esk2_0,k1_xboole_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_101]), c_0_45]), c_0_47])]), c_0_48])).
cnf(c_0_117, negated_conjecture, (~l1_lattices(esk2_0)|~m1_subset_1(k15_lattice3(esk2_0,k1_xboole_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_106]), c_0_45]), c_0_47])]), c_0_48])).
cnf(c_0_118, negated_conjecture, (~l1_lattices(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_33]), c_0_45])]), c_0_48])).
cnf(c_0_119, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_104]), c_0_45])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 120
# Proof object clause steps            : 77
# Proof object formula steps           : 43
# Proof object conjectures             : 32
# Proof object clause conjectures      : 29
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 32
# Proof object initial formulas used   : 21
# Proof object generating inferences   : 37
# Proof object simplifying inferences  : 108
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 67
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 66
# Processed clauses                    : 1472
# ...of these trivial                  : 12
# ...subsumed                          : 479
# ...remaining for further processing  : 981
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 125
# Backward-rewritten                   : 0
# Generated clauses                    : 15493
# ...of the previous two non-trivial   : 15148
# Contextual simplify-reflections      : 142
# Paramodulations                      : 15483
# Factorizations                       : 2
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 782
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 775
# Current number of unprocessed clauses: 13800
# ...number of literals in the above   : 66291
# Current number of archived formulas  : 0
# Current number of archived clauses   : 191
# Clause-clause subsumption calls (NU) : 67639
# Rec. Clause-clause subsumption calls : 9592
# Non-unit clause-clause subsumptions  : 746
# Unit Clause-clause subsumption calls : 413
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 517652

# -------------------------------------------------
# User time                : 0.239 s
# System time              : 0.003 s
# Total time               : 0.243 s
# Maximum resident set size: 31508 pages
