# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t52_relset_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(~(v1_xboole_0(X3))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))=>![X5]:(m1_subset_1(X5,X1)=>(r2_hidden(X5,k7_relset_1(X3,X1,X4,X2))<=>?[X6]:(m1_subset_1(X6,X3)&(r2_hidden(k4_tarski(X6,X5),X4)&r2_hidden(X6,X2))))))))), file('data/MPTP/proofs/2499226260254e9385f4074531d6b8b5.E_input', t52_relset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/2499226260254e9385f4074531d6b8b5.E_input', t3_subset)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/2499226260254e9385f4074531d6b8b5.E_input', d3_tarski)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/2499226260254e9385f4074531d6b8b5.E_input', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/2499226260254e9385f4074531d6b8b5.E_input', fc6_relat_1)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/2499226260254e9385f4074531d6b8b5.E_input', l54_zfmisc_1)).
fof(redefinition_k7_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k7_relset_1(X1,X2,X3,X4)=k9_relat_1(X3,X4)), file('data/MPTP/proofs/2499226260254e9385f4074531d6b8b5.E_input', redefinition_k7_relset_1)).
fof(t143_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k9_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k1_relat_1(X3))&(r2_hidden(k4_tarski(X4,X1),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/2499226260254e9385f4074531d6b8b5.E_input', t143_relat_1)).
fof(d13_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(X3=k9_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(k4_tarski(X5,X4),X1)&r2_hidden(X5,X2))))), file('data/MPTP/proofs/2499226260254e9385f4074531d6b8b5.E_input', d13_relat_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/2499226260254e9385f4074531d6b8b5.E_input', t1_subset)).
fof(c_0_10, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(~(v1_xboole_0(X3))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))=>![X5]:(m1_subset_1(X5,X1)=>(r2_hidden(X5,k7_relset_1(X3,X1,X4,X2))<=>?[X6]:(m1_subset_1(X6,X3)&(r2_hidden(k4_tarski(X6,X5),X4)&r2_hidden(X6,X2)))))))))), inference(assume_negation,[status(cth)],[t52_relset_1])).
fof(c_0_11, plain, ![X31, X32]:((~m1_subset_1(X31,k1_zfmisc_1(X32))|r1_tarski(X31,X32))&(~r1_tarski(X31,X32)|m1_subset_1(X31,k1_zfmisc_1(X32)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_12, negated_conjecture, ![X12]:(~v1_xboole_0(esk1_0)&(~v1_xboole_0(esk2_0)&(~v1_xboole_0(esk3_0)&(m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk1_0)))&(m1_subset_1(esk5_0,esk1_0)&((~r2_hidden(esk5_0,k7_relset_1(esk3_0,esk1_0,esk4_0,esk2_0))|(~m1_subset_1(X12,esk3_0)|(~r2_hidden(k4_tarski(X12,esk5_0),esk4_0)|~r2_hidden(X12,esk2_0))))&((m1_subset_1(esk6_0,esk3_0)|r2_hidden(esk5_0,k7_relset_1(esk3_0,esk1_0,esk4_0,esk2_0)))&((r2_hidden(k4_tarski(esk6_0,esk5_0),esk4_0)|r2_hidden(esk5_0,k7_relset_1(esk3_0,esk1_0,esk4_0,esk2_0)))&(r2_hidden(esk6_0,esk2_0)|r2_hidden(esk5_0,k7_relset_1(esk3_0,esk1_0,esk4_0,esk2_0))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_10])])])])])])).
fof(c_0_13, plain, ![X35, X36, X37, X38, X39]:((~r1_tarski(X35,X36)|(~r2_hidden(X37,X35)|r2_hidden(X37,X36)))&((r2_hidden(esk8_2(X38,X39),X38)|r1_tarski(X38,X39))&(~r2_hidden(esk8_2(X38,X39),X39)|r1_tarski(X38,X39)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_14, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_16, plain, ![X18, X19]:(~v1_relat_1(X18)|(~m1_subset_1(X19,k1_zfmisc_1(X18))|v1_relat_1(X19))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_17, plain, ![X33, X34]:v1_relat_1(k2_zfmisc_1(X33,X34)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_18, plain, ![X25, X26, X27, X28]:(((r2_hidden(X25,X27)|~r2_hidden(k4_tarski(X25,X26),k2_zfmisc_1(X27,X28)))&(r2_hidden(X26,X28)|~r2_hidden(k4_tarski(X25,X26),k2_zfmisc_1(X27,X28))))&(~r2_hidden(X25,X27)|~r2_hidden(X26,X28)|r2_hidden(k4_tarski(X25,X26),k2_zfmisc_1(X27,X28)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
cnf(c_0_19, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, negated_conjecture, (r1_tarski(esk4_0,k2_zfmisc_1(esk3_0,esk1_0))), inference(spm,[status(thm)],[c_0_14, c_0_15])).
fof(c_0_21, plain, ![X14, X15, X16, X17]:(~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15)))|k7_relset_1(X14,X15,X16,X17)=k9_relat_1(X16,X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_relset_1])])).
cnf(c_0_22, negated_conjecture, (~r2_hidden(esk5_0,k7_relset_1(esk3_0,esk1_0,esk4_0,esk2_0))|~m1_subset_1(X1,esk3_0)|~r2_hidden(k4_tarski(X1,esk5_0),esk4_0)|~r2_hidden(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23, negated_conjecture, (r2_hidden(esk6_0,esk2_0)|r2_hidden(esk5_0,k7_relset_1(esk3_0,esk1_0,esk4_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_24, plain, ![X20, X21, X22, X24]:(((r2_hidden(esk7_3(X20,X21,X22),k1_relat_1(X22))|~r2_hidden(X20,k9_relat_1(X22,X21))|~v1_relat_1(X22))&((r2_hidden(k4_tarski(esk7_3(X20,X21,X22),X20),X22)|~r2_hidden(X20,k9_relat_1(X22,X21))|~v1_relat_1(X22))&(r2_hidden(esk7_3(X20,X21,X22),X21)|~r2_hidden(X20,k9_relat_1(X22,X21))|~v1_relat_1(X22))))&(~r2_hidden(X24,k1_relat_1(X22))|(~r2_hidden(k4_tarski(X24,X20),X22)|~r2_hidden(X24,X21))|r2_hidden(X20,k9_relat_1(X22,X21))|~v1_relat_1(X22))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t143_relat_1])])])])])).
cnf(c_0_25, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(esk3_0,esk1_0))|~r2_hidden(X1,esk4_0)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
fof(c_0_29, plain, ![X41, X42, X43, X44, X46, X47, X48, X49, X51]:((((r2_hidden(k4_tarski(esk9_4(X41,X42,X43,X44),X44),X41)|~r2_hidden(X44,X43)|X43!=k9_relat_1(X41,X42)|~v1_relat_1(X41))&(r2_hidden(esk9_4(X41,X42,X43,X44),X42)|~r2_hidden(X44,X43)|X43!=k9_relat_1(X41,X42)|~v1_relat_1(X41)))&(~r2_hidden(k4_tarski(X47,X46),X41)|~r2_hidden(X47,X42)|r2_hidden(X46,X43)|X43!=k9_relat_1(X41,X42)|~v1_relat_1(X41)))&((~r2_hidden(esk10_3(X41,X48,X49),X49)|(~r2_hidden(k4_tarski(X51,esk10_3(X41,X48,X49)),X41)|~r2_hidden(X51,X48))|X49=k9_relat_1(X41,X48)|~v1_relat_1(X41))&((r2_hidden(k4_tarski(esk11_3(X41,X48,X49),esk10_3(X41,X48,X49)),X41)|r2_hidden(esk10_3(X41,X48,X49),X49)|X49=k9_relat_1(X41,X48)|~v1_relat_1(X41))&(r2_hidden(esk11_3(X41,X48,X49),X48)|r2_hidden(esk10_3(X41,X48,X49),X49)|X49=k9_relat_1(X41,X48)|~v1_relat_1(X41))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d13_relat_1])])])])])])).
cnf(c_0_30, plain, (k7_relset_1(X2,X3,X1,X4)=k9_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, negated_conjecture, (r2_hidden(esk6_0,esk2_0)|~r2_hidden(k4_tarski(X1,esk5_0),esk4_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk3_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_32, plain, (r2_hidden(k4_tarski(esk7_3(X1,X2,X3),X1),X3)|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, negated_conjecture, (v1_relat_1(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_15]), c_0_26])])).
fof(c_0_34, plain, ![X29, X30]:(~r2_hidden(X29,X30)|m1_subset_1(X29,X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_35, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(k4_tarski(X1,X2),esk4_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_36, plain, (r2_hidden(X2,X5)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,X4)|X5!=k9_relat_1(X3,X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, negated_conjecture, (r2_hidden(k4_tarski(esk6_0,esk5_0),esk4_0)|r2_hidden(esk5_0,k7_relset_1(esk3_0,esk1_0,esk4_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_38, negated_conjecture, (k7_relset_1(esk3_0,esk1_0,esk4_0,X1)=k9_relat_1(esk4_0,X1)), inference(spm,[status(thm)],[c_0_30, c_0_15])).
cnf(c_0_39, negated_conjecture, (r2_hidden(esk6_0,esk2_0)|~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk2_0)|~r2_hidden(esk5_0,k9_relat_1(esk4_0,X1))|~m1_subset_1(esk7_3(esk5_0,X1,esk4_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])])).
cnf(c_0_40, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_41, negated_conjecture, (r2_hidden(esk7_3(X1,X2,esk4_0),esk3_0)|~r2_hidden(X1,k9_relat_1(esk4_0,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_32]), c_0_33])])).
cnf(c_0_42, plain, (r2_hidden(X1,k9_relat_1(X2,X3))|~v1_relat_1(X2)|~r2_hidden(k4_tarski(X4,X1),X2)|~r2_hidden(X4,X3)), inference(er,[status(thm)],[c_0_36])).
cnf(c_0_43, negated_conjecture, (r2_hidden(k4_tarski(esk6_0,esk5_0),esk4_0)|r2_hidden(esk5_0,k9_relat_1(esk4_0,esk2_0))), inference(rw,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_44, negated_conjecture, (r2_hidden(esk6_0,esk2_0)|~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk2_0)|~r2_hidden(esk5_0,k9_relat_1(esk4_0,X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])).
cnf(c_0_45, plain, (r2_hidden(esk7_3(X1,X2,X3),X2)|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk5_0,k9_relat_1(esk4_0,esk2_0))|r2_hidden(esk6_0,esk2_0)), inference(rw,[status(thm)],[c_0_23, c_0_38])).
cnf(c_0_47, negated_conjecture, (r2_hidden(esk5_0,k9_relat_1(esk4_0,esk2_0))|r2_hidden(esk5_0,k9_relat_1(esk4_0,X1))|~r2_hidden(esk6_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_33])])).
cnf(c_0_48, negated_conjecture, (r2_hidden(esk6_0,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_33])]), c_0_46])).
cnf(c_0_49, negated_conjecture, (~r2_hidden(esk5_0,k9_relat_1(esk4_0,esk2_0))|~r2_hidden(k4_tarski(X1,esk5_0),esk4_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk3_0)), inference(rw,[status(thm)],[c_0_22, c_0_38])).
cnf(c_0_50, negated_conjecture, (r2_hidden(esk5_0,k9_relat_1(esk4_0,esk2_0))), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_51, negated_conjecture, (~r2_hidden(k4_tarski(X1,esk5_0),esk4_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_50])])).
cnf(c_0_52, negated_conjecture, (~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk2_0)|~r2_hidden(esk5_0,k9_relat_1(esk4_0,X1))|~m1_subset_1(esk7_3(esk5_0,X1,esk4_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_32]), c_0_33])])).
cnf(c_0_53, negated_conjecture, (~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk2_0)|~r2_hidden(esk5_0,k9_relat_1(esk4_0,X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_40]), c_0_41])).
cnf(c_0_54, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_45]), c_0_50]), c_0_33])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 55
# Proof object clause steps            : 34
# Proof object formula steps           : 21
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 24
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 31
# Processed clauses                    : 146
# ...of these trivial                  : 4
# ...subsumed                          : 18
# ...remaining for further processing  : 124
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 9
# Backward-rewritten                   : 17
# Generated clauses                    : 232
# ...of the previous two non-trivial   : 209
# Contextual simplify-reflections      : 6
# Paramodulations                      : 229
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 65
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 48
# Current number of unprocessed clauses: 123
# ...number of literals in the above   : 456
# Current number of archived formulas  : 0
# Current number of archived clauses   : 56
# Clause-clause subsumption calls (NU) : 1584
# Rec. Clause-clause subsumption calls : 1112
# Non-unit clause-clause subsumptions  : 28
# Unit Clause-clause subsumption calls : 52
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12
# BW rewrite match successes           : 7
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6146

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.000 s
# Total time               : 0.013 s
# Maximum resident set size: 28720 pages
