# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t44_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_tarski(X1),k1_enumset1(X2,X3,X4)), file('data/MPTP/proofs/249bbec037434dc98d3a95afe3abb27a.E_input', t44_enumset1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/249bbec037434dc98d3a95afe3abb27a.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/249bbec037434dc98d3a95afe3abb27a.E_input', t70_enumset1)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/249bbec037434dc98d3a95afe3abb27a.E_input', t98_xboole_1)).
fof(l62_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k1_enumset1(X1,X2,X3),k1_enumset1(X4,X5,X6)), file('data/MPTP/proofs/249bbec037434dc98d3a95afe3abb27a.E_input', l62_enumset1)).
fof(t48_enumset1, axiom, ![X1, X2, X3, X4, X5]:k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_tarski(X1,X2),k1_enumset1(X3,X4,X5)), file('data/MPTP/proofs/249bbec037434dc98d3a95afe3abb27a.E_input', t48_enumset1)).
fof(t52_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k2_tarski(X1,X2),k2_enumset1(X3,X4,X5,X6)), file('data/MPTP/proofs/249bbec037434dc98d3a95afe3abb27a.E_input', t52_enumset1)).
fof(t63_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_tarski(X1,X2),k4_enumset1(X3,X4,X5,X6,X7,X8)), file('data/MPTP/proofs/249bbec037434dc98d3a95afe3abb27a.E_input', t63_enumset1)).
fof(t47_enumset1, axiom, ![X1, X2, X3, X4, X5]:k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k1_tarski(X1),k2_enumset1(X2,X3,X4,X5)), file('data/MPTP/proofs/249bbec037434dc98d3a95afe3abb27a.E_input', t47_enumset1)).
fof(t64_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k1_enumset1(X1,X2,X3),k3_enumset1(X4,X5,X6,X7,X8)), file('data/MPTP/proofs/249bbec037434dc98d3a95afe3abb27a.E_input', t64_enumset1)).
fof(t43_enumset1, axiom, ![X1, X2, X3]:k1_enumset1(X1,X2,X3)=k2_xboole_0(k2_tarski(X1,X2),k1_tarski(X3)), file('data/MPTP/proofs/249bbec037434dc98d3a95afe3abb27a.E_input', t43_enumset1)).
fof(t71_enumset1, conjecture, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/249bbec037434dc98d3a95afe3abb27a.E_input', t71_enumset1)).
fof(t55_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_tarski(X6)), file('data/MPTP/proofs/249bbec037434dc98d3a95afe3abb27a.E_input', t55_enumset1)).
fof(c_0_13, plain, ![X9, X10, X11, X12]:k2_enumset1(X9,X10,X11,X12)=k2_xboole_0(k1_tarski(X9),k1_enumset1(X10,X11,X12)), inference(variable_rename,[status(thm)],[t44_enumset1])).
fof(c_0_14, plain, ![X13]:k2_tarski(X13,X13)=k1_tarski(X13), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_15, plain, ![X14, X15]:k1_enumset1(X14,X14,X15)=k2_tarski(X14,X15), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_16, plain, ![X16, X17]:k2_xboole_0(X16,X17)=k5_xboole_0(X16,k4_xboole_0(X17,X16)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_17, plain, ![X18, X19, X20, X21, X22, X23]:k4_enumset1(X18,X19,X20,X21,X22,X23)=k2_xboole_0(k1_enumset1(X18,X19,X20),k1_enumset1(X21,X22,X23)), inference(variable_rename,[status(thm)],[l62_enumset1])).
fof(c_0_18, plain, ![X24, X25, X26, X27, X28]:k3_enumset1(X24,X25,X26,X27,X28)=k2_xboole_0(k2_tarski(X24,X25),k1_enumset1(X26,X27,X28)), inference(variable_rename,[status(thm)],[t48_enumset1])).
fof(c_0_19, plain, ![X29, X30, X31, X32, X33, X34]:k4_enumset1(X29,X30,X31,X32,X33,X34)=k2_xboole_0(k2_tarski(X29,X30),k2_enumset1(X31,X32,X33,X34)), inference(variable_rename,[status(thm)],[t52_enumset1])).
cnf(c_0_20, plain, (k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_tarski(X1),k1_enumset1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k1_enumset1(X1,X2,X3),k1_enumset1(X4,X5,X6))), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_25, plain, ![X35, X36, X37, X38, X39, X40, X41, X42]:k6_enumset1(X35,X36,X37,X38,X39,X40,X41,X42)=k2_xboole_0(k2_tarski(X35,X36),k4_enumset1(X37,X38,X39,X40,X41,X42)), inference(variable_rename,[status(thm)],[t63_enumset1])).
fof(c_0_26, plain, ![X43, X44, X45, X46, X47]:k3_enumset1(X43,X44,X45,X46,X47)=k2_xboole_0(k1_tarski(X43),k2_enumset1(X44,X45,X46,X47)), inference(variable_rename,[status(thm)],[t47_enumset1])).
cnf(c_0_27, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_tarski(X1,X2),k1_enumset1(X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, plain, (k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k2_tarski(X1,X2),k2_enumset1(X3,X4,X5,X6))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (k2_enumset1(X1,X2,X3,X4)=k5_xboole_0(k1_enumset1(X1,X1,X1),k4_xboole_0(k1_enumset1(X2,X3,X4),k1_enumset1(X1,X1,X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23])).
cnf(c_0_30, plain, (k4_enumset1(X1,X2,X3,X4,X5,X6)=k5_xboole_0(k1_enumset1(X1,X2,X3),k4_xboole_0(k1_enumset1(X4,X5,X6),k1_enumset1(X1,X2,X3)))), inference(rw,[status(thm)],[c_0_24, c_0_23])).
cnf(c_0_31, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_tarski(X1,X2),k4_enumset1(X3,X4,X5,X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_32, plain, ![X48, X49, X50, X51, X52, X53, X54, X55]:k6_enumset1(X48,X49,X50,X51,X52,X53,X54,X55)=k2_xboole_0(k1_enumset1(X48,X49,X50),k3_enumset1(X51,X52,X53,X54,X55)), inference(variable_rename,[status(thm)],[t64_enumset1])).
fof(c_0_33, plain, ![X56, X57, X58]:k1_enumset1(X56,X57,X58)=k2_xboole_0(k2_tarski(X56,X57),k1_tarski(X58)), inference(variable_rename,[status(thm)],[t43_enumset1])).
cnf(c_0_34, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k1_tarski(X1),k2_enumset1(X2,X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k5_xboole_0(k1_enumset1(X1,X1,X2),k4_xboole_0(k1_enumset1(X3,X4,X5),k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_22]), c_0_23])).
cnf(c_0_36, plain, (k5_xboole_0(k1_enumset1(X1,X2,X3),k4_xboole_0(k1_enumset1(X4,X5,X6),k1_enumset1(X1,X2,X3)))=k5_xboole_0(k1_enumset1(X1,X1,X2),k4_xboole_0(k5_xboole_0(k1_enumset1(X3,X3,X3),k4_xboole_0(k1_enumset1(X4,X5,X6),k1_enumset1(X3,X3,X3))),k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_22]), c_0_23]), c_0_29]), c_0_30])).
cnf(c_0_37, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k5_xboole_0(k1_enumset1(X1,X1,X2),k4_xboole_0(k5_xboole_0(k1_enumset1(X3,X4,X5),k4_xboole_0(k1_enumset1(X6,X7,X8),k1_enumset1(X3,X4,X5))),k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_22]), c_0_23]), c_0_30])).
cnf(c_0_38, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k1_enumset1(X1,X2,X3),k3_enumset1(X4,X5,X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39, plain, (k1_enumset1(X1,X2,X3)=k2_xboole_0(k2_tarski(X1,X2),k1_tarski(X3))), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_40, negated_conjecture, ~(![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(assume_negation,[status(cth)],[t71_enumset1])).
fof(c_0_41, plain, ![X62, X63, X64, X65, X66, X67]:k4_enumset1(X62,X63,X64,X65,X66,X67)=k2_xboole_0(k3_enumset1(X62,X63,X64,X65,X66),k1_tarski(X67)), inference(variable_rename,[status(thm)],[t55_enumset1])).
cnf(c_0_42, plain, (k5_xboole_0(k1_enumset1(X1,X1,X2),k4_xboole_0(k1_enumset1(X3,X4,X5),k1_enumset1(X1,X1,X2)))=k5_xboole_0(k1_enumset1(X1,X1,X1),k4_xboole_0(k5_xboole_0(k1_enumset1(X2,X2,X2),k4_xboole_0(k1_enumset1(X3,X4,X5),k1_enumset1(X2,X2,X2))),k1_enumset1(X1,X1,X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_21]), c_0_22]), c_0_23]), c_0_29]), c_0_35])).
cnf(c_0_43, plain, (k6_enumset1(X1,X2,X3,X3,X3,X4,X5,X6)=k5_xboole_0(k1_enumset1(X1,X2,X3),k4_xboole_0(k1_enumset1(X4,X5,X6),k1_enumset1(X1,X2,X3)))), inference(rw,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_44, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k5_xboole_0(k1_enumset1(X1,X2,X3),k4_xboole_0(k5_xboole_0(k1_enumset1(X4,X4,X5),k4_xboole_0(k1_enumset1(X6,X7,X8),k1_enumset1(X4,X4,X5))),k1_enumset1(X1,X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_23]), c_0_35])).
cnf(c_0_45, plain, (k1_enumset1(X1,X2,X3)=k5_xboole_0(k1_enumset1(X1,X1,X2),k4_xboole_0(k1_enumset1(X3,X3,X3),k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_21]), c_0_22]), c_0_22]), c_0_23])).
fof(c_0_46, negated_conjecture, k2_enumset1(esk1_0,esk1_0,esk2_0,esk3_0)!=k1_enumset1(esk1_0,esk2_0,esk3_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_40])])])).
cnf(c_0_47, plain, (k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_tarski(X6))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_48, plain, (k5_xboole_0(k1_enumset1(X1,X1,X1),k4_xboole_0(k6_enumset1(X2,X2,X2,X2,X2,X3,X4,X5),k1_enumset1(X1,X1,X1)))=k5_xboole_0(k1_enumset1(X1,X1,X2),k4_xboole_0(k1_enumset1(X3,X4,X5),k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_49, plain, (k5_xboole_0(k1_enumset1(X1,X2,X3),k4_xboole_0(k1_enumset1(X4,X5,X6),k1_enumset1(X1,X2,X3)))=k6_enumset1(X1,X2,X3,X4,X5,X6,X6,X6)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_50, negated_conjecture, (k2_enumset1(esk1_0,esk1_0,esk2_0,esk3_0)!=k1_enumset1(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_51, plain, (k5_xboole_0(k1_enumset1(X1,X2,X3),k4_xboole_0(k1_enumset1(X4,X5,X6),k1_enumset1(X1,X2,X3)))=k5_xboole_0(k5_xboole_0(k1_enumset1(X1,X1,X2),k4_xboole_0(k1_enumset1(X3,X4,X5),k1_enumset1(X1,X1,X2))),k4_xboole_0(k1_enumset1(X6,X6,X6),k5_xboole_0(k1_enumset1(X1,X1,X2),k4_xboole_0(k1_enumset1(X3,X4,X5),k1_enumset1(X1,X1,X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_21]), c_0_22]), c_0_23]), c_0_35]), c_0_35]), c_0_30])).
cnf(c_0_52, plain, (k6_enumset1(X1,X2,X3,X3,X4,X5,X6,X7)=k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)), inference(spm,[status(thm)],[c_0_44, c_0_37])).
cnf(c_0_53, plain, (k6_enumset1(X1,X1,X2,X2,X2,X2,X2,X3)=k1_enumset1(X1,X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_37]), c_0_45])).
cnf(c_0_54, plain, (k5_xboole_0(k1_enumset1(X1,X2,X3),k4_xboole_0(k6_enumset1(X4,X4,X5,X5,X5,X6,X7,X8),k1_enumset1(X1,X2,X3)))=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_44, c_0_43])).
cnf(c_0_55, negated_conjecture, (k5_xboole_0(k1_enumset1(esk1_0,esk1_0,esk1_0),k4_xboole_0(k1_enumset1(esk1_0,esk2_0,esk3_0),k1_enumset1(esk1_0,esk1_0,esk1_0)))!=k1_enumset1(esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_50, c_0_29])).
cnf(c_0_56, plain, (k5_xboole_0(k6_enumset1(X1,X1,X2,X2,X2,X3,X4,X5),k4_xboole_0(k1_enumset1(X6,X6,X6),k6_enumset1(X1,X1,X2,X2,X2,X3,X4,X5)))=k5_xboole_0(k1_enumset1(X1,X2,X3),k4_xboole_0(k1_enumset1(X4,X5,X6),k1_enumset1(X1,X2,X3)))), inference(spm,[status(thm)],[c_0_51, c_0_43])).
cnf(c_0_57, plain, (k6_enumset1(X1,X1,X1,X2,X2,X2,X2,X3)=k1_enumset1(X1,X2,X3)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_58, plain, (k5_xboole_0(k1_enumset1(X1,X1,X2),k4_xboole_0(k1_enumset1(X3,X4,X5),k1_enumset1(X1,X1,X2)))=k6_enumset1(X1,X1,X1,X2,X2,X3,X4,X5)), inference(rw,[status(thm)],[c_0_48, c_0_54])).
cnf(c_0_59, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)!=k1_enumset1(esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_55, c_0_43])).
cnf(c_0_60, plain, (k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_45]), c_0_58])).
cnf(c_0_61, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_60])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 62
# Proof object clause steps            : 35
# Proof object formula steps           : 27
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 41
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 13
# Removed in clause preprocessing      : 6
# Initial clauses in saturation        : 7
# Processed clauses                    : 102
# ...of these trivial                  : 24
# ...subsumed                          : 35
# ...remaining for further processing  : 43
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 9
# Generated clauses                    : 354
# ...of the previous two non-trivial   : 226
# Contextual simplify-reflections      : 0
# Paramodulations                      : 354
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 27
#    Positive orientable unit clauses  : 21
#    Positive unorientable unit clauses: 6
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 137
# ...number of literals in the above   : 137
# Current number of archived formulas  : 0
# Current number of archived clauses   : 22
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 23
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 411
# BW rewrite match successes           : 44
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3818

# -------------------------------------------------
# User time                : 0.005 s
# System time              : 0.005 s
# Total time               : 0.011 s
# Maximum resident set size: 28792 pages
