# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t78_mcart_1, axiom, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&?[X6, X7, X8, X9]:(X5=k4_mcart_1(X6,X7,X8,X9)&~((k8_mcart_1(X1,X2,X3,X4,X5)=X6&(k9_mcart_1(X1,X2,X3,X4,X5)=X7&(k10_mcart_1(X1,X2,X3,X4,X5)=X8&k11_mcart_1(X1,X2,X3,X4,X5)=X9))))))))))), file('data/MPTP/proofs/24cc2c9e27614853b6b4ff8ce0923750.E_input', t78_mcart_1)).
fof(d4_zfmisc_1, axiom, ![X1, X2, X3, X4]:k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4), file('data/MPTP/proofs/24cc2c9e27614853b6b4ff8ce0923750.E_input', d4_zfmisc_1)).
fof(t81_mcart_1, conjecture, ![X1, X2, X3, X4, X5, X6]:(m1_subset_1(X6,k4_zfmisc_1(X1,X2,X3,X4))=>(![X7]:(m1_subset_1(X7,X1)=>![X8]:(m1_subset_1(X8,X2)=>![X9]:(m1_subset_1(X9,X3)=>![X10]:(m1_subset_1(X10,X4)=>(X6=k4_mcart_1(X7,X8,X9,X10)=>X5=X9)))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|(X4=k1_xboole_0|X5=k10_mcart_1(X1,X2,X3,X4,X6))))))), file('data/MPTP/proofs/24cc2c9e27614853b6b4ff8ce0923750.E_input', t81_mcart_1)).
fof(l68_mcart_1, axiom, ![X1, X2, X3, X4]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&?[X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))&![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>![X9]:(m1_subset_1(X9,X4)=>~(X5=k4_mcart_1(X6,X7,X8,X9)))))))))))), file('data/MPTP/proofs/24cc2c9e27614853b6b4ff8ce0923750.E_input', l68_mcart_1)).
fof(t61_mcart_1, axiom, ![X1, X2, X3, X4]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&~(![X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>(k8_mcart_1(X1,X2,X3,X4,X5)=k1_mcart_1(k1_mcart_1(k1_mcart_1(X5)))&(k9_mcart_1(X1,X2,X3,X4,X5)=k2_mcart_1(k1_mcart_1(k1_mcart_1(X5)))&(k10_mcart_1(X1,X2,X3,X4,X5)=k2_mcart_1(k1_mcart_1(X5))&k11_mcart_1(X1,X2,X3,X4,X5)=k2_mcart_1(X5))))))))))), file('data/MPTP/proofs/24cc2c9e27614853b6b4ff8ce0923750.E_input', t61_mcart_1)).
fof(c_0_5, plain, ![X39, X40, X41, X42, X43, X44, X45, X46, X47]:((k8_mcart_1(X39,X40,X41,X42,X43)=X44|X43!=k4_mcart_1(X44,X45,X46,X47)|X42=k1_xboole_0|X41=k1_xboole_0|X40=k1_xboole_0|X39=k1_xboole_0|~m1_subset_1(X43,k4_zfmisc_1(X39,X40,X41,X42)))&((k9_mcart_1(X39,X40,X41,X42,X43)=X45|X43!=k4_mcart_1(X44,X45,X46,X47)|X42=k1_xboole_0|X41=k1_xboole_0|X40=k1_xboole_0|X39=k1_xboole_0|~m1_subset_1(X43,k4_zfmisc_1(X39,X40,X41,X42)))&((k10_mcart_1(X39,X40,X41,X42,X43)=X46|X43!=k4_mcart_1(X44,X45,X46,X47)|X42=k1_xboole_0|X41=k1_xboole_0|X40=k1_xboole_0|X39=k1_xboole_0|~m1_subset_1(X43,k4_zfmisc_1(X39,X40,X41,X42)))&(k11_mcart_1(X39,X40,X41,X42,X43)=X47|X43!=k4_mcart_1(X44,X45,X46,X47)|X42=k1_xboole_0|X41=k1_xboole_0|X40=k1_xboole_0|X39=k1_xboole_0|~m1_subset_1(X43,k4_zfmisc_1(X39,X40,X41,X42)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t78_mcart_1])])])])])).
fof(c_0_6, plain, ![X21, X22, X23, X24]:k4_zfmisc_1(X21,X22,X23,X24)=k2_zfmisc_1(k3_zfmisc_1(X21,X22,X23),X24), inference(variable_rename,[status(thm)],[d4_zfmisc_1])).
fof(c_0_7, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6]:(m1_subset_1(X6,k4_zfmisc_1(X1,X2,X3,X4))=>(![X7]:(m1_subset_1(X7,X1)=>![X8]:(m1_subset_1(X8,X2)=>![X9]:(m1_subset_1(X9,X3)=>![X10]:(m1_subset_1(X10,X4)=>(X6=k4_mcart_1(X7,X8,X9,X10)=>X5=X9)))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|(X4=k1_xboole_0|X5=k10_mcart_1(X1,X2,X3,X4,X6)))))))), inference(assume_negation,[status(cth)],[t81_mcart_1])).
cnf(c_0_8, plain, (k10_mcart_1(X1,X2,X3,X4,X5)=X6|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X5!=k4_mcart_1(X7,X8,X6,X9)|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_10, plain, ![X30, X31, X32, X33, X34]:((m1_subset_1(esk7_5(X30,X31,X32,X33,X34),X30)|~m1_subset_1(X34,k4_zfmisc_1(X30,X31,X32,X33))|X33=k1_xboole_0|X32=k1_xboole_0|X31=k1_xboole_0|X30=k1_xboole_0)&((m1_subset_1(esk8_5(X30,X31,X32,X33,X34),X31)|~m1_subset_1(X34,k4_zfmisc_1(X30,X31,X32,X33))|X33=k1_xboole_0|X32=k1_xboole_0|X31=k1_xboole_0|X30=k1_xboole_0)&((m1_subset_1(esk9_5(X30,X31,X32,X33,X34),X32)|~m1_subset_1(X34,k4_zfmisc_1(X30,X31,X32,X33))|X33=k1_xboole_0|X32=k1_xboole_0|X31=k1_xboole_0|X30=k1_xboole_0)&((m1_subset_1(esk10_5(X30,X31,X32,X33,X34),X33)|~m1_subset_1(X34,k4_zfmisc_1(X30,X31,X32,X33))|X33=k1_xboole_0|X32=k1_xboole_0|X31=k1_xboole_0|X30=k1_xboole_0)&(X34=k4_mcart_1(esk7_5(X30,X31,X32,X33,X34),esk8_5(X30,X31,X32,X33,X34),esk9_5(X30,X31,X32,X33,X34),esk10_5(X30,X31,X32,X33,X34))|~m1_subset_1(X34,k4_zfmisc_1(X30,X31,X32,X33))|X33=k1_xboole_0|X32=k1_xboole_0|X31=k1_xboole_0|X30=k1_xboole_0))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l68_mcart_1])])])])])])])).
fof(c_0_11, plain, ![X25, X26, X27, X28, X29]:((k8_mcart_1(X25,X26,X27,X28,X29)=k1_mcart_1(k1_mcart_1(k1_mcart_1(X29)))|~m1_subset_1(X29,k4_zfmisc_1(X25,X26,X27,X28))|X28=k1_xboole_0|X27=k1_xboole_0|X26=k1_xboole_0|X25=k1_xboole_0)&((k9_mcart_1(X25,X26,X27,X28,X29)=k2_mcart_1(k1_mcart_1(k1_mcart_1(X29)))|~m1_subset_1(X29,k4_zfmisc_1(X25,X26,X27,X28))|X28=k1_xboole_0|X27=k1_xboole_0|X26=k1_xboole_0|X25=k1_xboole_0)&((k10_mcart_1(X25,X26,X27,X28,X29)=k2_mcart_1(k1_mcart_1(X29))|~m1_subset_1(X29,k4_zfmisc_1(X25,X26,X27,X28))|X28=k1_xboole_0|X27=k1_xboole_0|X26=k1_xboole_0|X25=k1_xboole_0)&(k11_mcart_1(X25,X26,X27,X28,X29)=k2_mcart_1(X29)|~m1_subset_1(X29,k4_zfmisc_1(X25,X26,X27,X28))|X28=k1_xboole_0|X27=k1_xboole_0|X26=k1_xboole_0|X25=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t61_mcart_1])])])])])).
fof(c_0_12, negated_conjecture, ![X17, X18, X19, X20]:(m1_subset_1(esk6_0,k4_zfmisc_1(esk1_0,esk2_0,esk3_0,esk4_0))&((~m1_subset_1(X17,esk1_0)|(~m1_subset_1(X18,esk2_0)|(~m1_subset_1(X19,esk3_0)|(~m1_subset_1(X20,esk4_0)|(esk6_0!=k4_mcart_1(X17,X18,X19,X20)|esk5_0=X19)))))&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&(esk4_0!=k1_xboole_0&esk5_0!=k10_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])).
cnf(c_0_13, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k10_mcart_1(X1,X2,X3,X4,X5)=X6|X5!=k4_mcart_1(X7,X8,X6,X9)|~m1_subset_1(X5,k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4))), inference(rw,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_14, plain, (X1=k4_mcart_1(esk7_5(X2,X3,X4,X5,X1),esk8_5(X2,X3,X4,X5,X1),esk9_5(X2,X3,X4,X5,X1),esk10_5(X2,X3,X4,X5,X1))|X5=k1_xboole_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|~m1_subset_1(X1,k4_zfmisc_1(X2,X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, plain, (k10_mcart_1(X1,X2,X3,X4,X5)=k2_mcart_1(k1_mcart_1(X5))|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, negated_conjecture, (m1_subset_1(esk6_0,k4_zfmisc_1(esk1_0,esk2_0,esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, plain, (k10_mcart_1(X1,X2,X3,X4,k4_mcart_1(X5,X6,X7,X8))=X7|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(k4_mcart_1(X5,X6,X7,X8),k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4))), inference(er,[status(thm)],[c_0_13])).
cnf(c_0_18, plain, (X5=k1_xboole_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k4_mcart_1(esk7_5(X2,X3,X4,X5,X1),esk8_5(X2,X3,X4,X5,X1),esk9_5(X2,X3,X4,X5,X1),esk10_5(X2,X3,X4,X5,X1))|~m1_subset_1(X1,k2_zfmisc_1(k3_zfmisc_1(X2,X3,X4),X5))), inference(rw,[status(thm)],[c_0_14, c_0_9])).
cnf(c_0_19, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k10_mcart_1(X1,X2,X3,X4,X5)=k2_mcart_1(k1_mcart_1(X5))|~m1_subset_1(X5,k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4))), inference(rw,[status(thm)],[c_0_15, c_0_9])).
cnf(c_0_20, negated_conjecture, (m1_subset_1(esk6_0,k2_zfmisc_1(k3_zfmisc_1(esk1_0,esk2_0,esk3_0),esk4_0))), inference(rw,[status(thm)],[c_0_16, c_0_9])).
cnf(c_0_21, negated_conjecture, (esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_24, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_25, plain, (k10_mcart_1(X1,X2,X3,X4,X5)=esk9_5(X6,X7,X8,X9,X5)|X6=k1_xboole_0|X7=k1_xboole_0|X8=k1_xboole_0|X9=k1_xboole_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4))|~m1_subset_1(X5,k2_zfmisc_1(k3_zfmisc_1(X6,X7,X8),X9))), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_26, negated_conjecture, (k10_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)=k2_mcart_1(k1_mcart_1(esk6_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24])).
cnf(c_0_27, negated_conjecture, (esk9_5(X1,X2,X3,X4,esk6_0)=k2_mcart_1(k1_mcart_1(esk6_0))|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk6_0,k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_20]), c_0_26]), c_0_21]), c_0_22]), c_0_23]), c_0_24])).
cnf(c_0_28, plain, (m1_subset_1(esk9_5(X1,X2,X3,X4,X5),X3)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_29, negated_conjecture, (esk9_5(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)=k2_mcart_1(k1_mcart_1(esk6_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24])).
cnf(c_0_30, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|m1_subset_1(esk9_5(X1,X2,X3,X4,X5),X3)|~m1_subset_1(X5,k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4))), inference(rw,[status(thm)],[c_0_28, c_0_9])).
cnf(c_0_31, negated_conjecture, (esk5_0!=k10_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_32, negated_conjecture, (esk5_0=X3|~m1_subset_1(X1,esk1_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X4,esk4_0)|esk6_0!=k4_mcart_1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_33, negated_conjecture, (k4_mcart_1(esk7_5(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0),esk8_5(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0),k2_mcart_1(k1_mcart_1(esk6_0)),esk10_5(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))=esk6_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_29]), c_0_20])]), c_0_21]), c_0_22]), c_0_23]), c_0_24])).
cnf(c_0_34, negated_conjecture, (m1_subset_1(k2_mcart_1(k1_mcart_1(esk6_0)),esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_29]), c_0_20])]), c_0_21]), c_0_22]), c_0_23]), c_0_24])).
cnf(c_0_35, negated_conjecture, (k2_mcart_1(k1_mcart_1(esk6_0))!=esk5_0), inference(rw,[status(thm)],[c_0_31, c_0_26])).
cnf(c_0_36, plain, (m1_subset_1(esk10_5(X1,X2,X3,X4,X5),X4)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_37, negated_conjecture, (~m1_subset_1(esk10_5(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0),esk4_0)|~m1_subset_1(esk8_5(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0),esk2_0)|~m1_subset_1(esk7_5(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_38, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|m1_subset_1(esk10_5(X1,X2,X3,X4,X5),X4)|~m1_subset_1(X5,k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4))), inference(rw,[status(thm)],[c_0_36, c_0_9])).
cnf(c_0_39, plain, (m1_subset_1(esk8_5(X1,X2,X3,X4,X5),X2)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_40, negated_conjecture, (~m1_subset_1(esk8_5(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0),esk2_0)|~m1_subset_1(esk7_5(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_20])]), c_0_21]), c_0_22]), c_0_23]), c_0_24])).
cnf(c_0_41, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|m1_subset_1(esk8_5(X1,X2,X3,X4,X5),X2)|~m1_subset_1(X5,k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4))), inference(rw,[status(thm)],[c_0_39, c_0_9])).
cnf(c_0_42, plain, (m1_subset_1(esk7_5(X1,X2,X3,X4,X5),X1)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_43, negated_conjecture, (~m1_subset_1(esk7_5(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_20])]), c_0_21]), c_0_22]), c_0_23]), c_0_24])).
cnf(c_0_44, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|m1_subset_1(esk7_5(X1,X2,X3,X4,X5),X1)|~m1_subset_1(X5,k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4))), inference(rw,[status(thm)],[c_0_42, c_0_9])).
cnf(c_0_45, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_20])]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 46
# Proof object clause steps            : 35
# Proof object formula steps           : 11
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 56
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 20
# Processed clauses                    : 58
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 58
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 1
# Generated clauses                    : 46
# ...of the previous two non-trivial   : 45
# Contextual simplify-reflections      : 0
# Paramodulations                      : 41
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 32
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 18
# Current number of unprocessed clauses: 27
# ...number of literals in the above   : 277
# Current number of archived formulas  : 0
# Current number of archived clauses   : 23
# Clause-clause subsumption calls (NU) : 46
# Rec. Clause-clause subsumption calls : 10
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 16
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3283

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28660 pages
