# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/24e5f60894954136a43e6cffc44b1ccf.E_input', d2_subset_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/24e5f60894954136a43e6cffc44b1ccf.E_input', d1_xboole_0)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/24e5f60894954136a43e6cffc44b1ccf.E_input', d8_lattice3)).
fof(d4_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>(v1_yellow_0(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&r1_lattice3(X1,u1_struct_0(X1),X2)))), file('data/MPTP/proofs/24e5f60894954136a43e6cffc44b1ccf.E_input', d4_yellow_0)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/24e5f60894954136a43e6cffc44b1ccf.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/24e5f60894954136a43e6cffc44b1ccf.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/24e5f60894954136a43e6cffc44b1ccf.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/24e5f60894954136a43e6cffc44b1ccf.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/24e5f60894954136a43e6cffc44b1ccf.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/24e5f60894954136a43e6cffc44b1ccf.E_input', t75_enumset1)).
fof(t44_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,k3_yellow_0(X1),X2))), file('data/MPTP/proofs/24e5f60894954136a43e6cffc44b1ccf.E_input', t44_yellow_0)).
fof(t8_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4)))&(((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4))=>r1_lattice3(X1,k2_tarski(X3,X4),X2))&((r2_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2)))&((r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2))=>r2_lattice3(X1,k2_tarski(X3,X4),X2))))))))), file('data/MPTP/proofs/24e5f60894954136a43e6cffc44b1ccf.E_input', t8_yellow_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/24e5f60894954136a43e6cffc44b1ccf.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/24e5f60894954136a43e6cffc44b1ccf.E_input', dt_l1_orders_2)).
fof(t30_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4)))))&((r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))=>(X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3)))))), file('data/MPTP/proofs/24e5f60894954136a43e6cffc44b1ccf.E_input', t30_yellow_0)).
fof(d11_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0)), file('data/MPTP/proofs/24e5f60894954136a43e6cffc44b1ccf.E_input', d11_yellow_0)).
fof(t6_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_lattice3(X1,k1_xboole_0,X2)&r1_lattice3(X1,k1_xboole_0,X2)))), file('data/MPTP/proofs/24e5f60894954136a43e6cffc44b1ccf.E_input', t6_yellow_0)).
fof(c_0_17, plain, ![X8, X9]:(((~m1_subset_1(X9,X8)|r2_hidden(X9,X8)|v1_xboole_0(X8))&(~r2_hidden(X9,X8)|m1_subset_1(X9,X8)|v1_xboole_0(X8)))&((~m1_subset_1(X9,X8)|v1_xboole_0(X9)|~v1_xboole_0(X8))&(~v1_xboole_0(X9)|m1_subset_1(X9,X8)|~v1_xboole_0(X8)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_18, plain, ![X10, X11, X12]:((~v1_xboole_0(X10)|~r2_hidden(X11,X10))&(r2_hidden(esk1_1(X12),X12)|v1_xboole_0(X12))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_19, plain, ![X14, X15, X16, X17]:((~r1_lattice3(X14,X15,X16)|(~m1_subset_1(X17,u1_struct_0(X14))|(~r2_hidden(X17,X15)|r1_orders_2(X14,X16,X17)))|~m1_subset_1(X16,u1_struct_0(X14))|~l1_orders_2(X14))&((m1_subset_1(esk2_3(X14,X15,X16),u1_struct_0(X14))|r1_lattice3(X14,X15,X16)|~m1_subset_1(X16,u1_struct_0(X14))|~l1_orders_2(X14))&((r2_hidden(esk2_3(X14,X15,X16),X15)|r1_lattice3(X14,X15,X16)|~m1_subset_1(X16,u1_struct_0(X14))|~l1_orders_2(X14))&(~r1_orders_2(X14,X16,esk2_3(X14,X15,X16))|r1_lattice3(X14,X15,X16)|~m1_subset_1(X16,u1_struct_0(X14))|~l1_orders_2(X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
fof(c_0_20, plain, ![X19, X21]:(((m1_subset_1(esk3_1(X19),u1_struct_0(X19))|~v1_yellow_0(X19)|~l1_orders_2(X19))&(r1_lattice3(X19,u1_struct_0(X19),esk3_1(X19))|~v1_yellow_0(X19)|~l1_orders_2(X19)))&(~m1_subset_1(X21,u1_struct_0(X19))|~r1_lattice3(X19,u1_struct_0(X19),X21)|v1_yellow_0(X19)|~l1_orders_2(X19))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_yellow_0])])])])])).
cnf(c_0_21, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_23, plain, ![X4, X3, X2, X1]:(epred1_4(X1,X2,X3,X4)<=>((r1_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4)))&(((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4))=>r1_lattice3(X1,k2_tarski(X3,X4),X2))&((r2_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2)))&((r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2))=>r2_lattice3(X1,k2_tarski(X3,X4),X2)))))), introduced(definition)).
cnf(c_0_24, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (r1_lattice3(X1,u1_struct_0(X1),esk3_1(X1))|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27, plain, (m1_subset_1(esk3_1(X1),u1_struct_0(X1))|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_28, plain, ![X4, X3, X2, X1]:(epred1_4(X1,X2,X3,X4)=>((r1_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4)))&(((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4))=>r1_lattice3(X1,k2_tarski(X3,X4),X2))&((r2_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2)))&((r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2))=>r2_lattice3(X1,k2_tarski(X3,X4),X2)))))), inference(split_equiv,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk2_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_30, plain, (r1_orders_2(X1,esk3_1(X1),X2)|~v1_yellow_0(X1)|~l1_orders_2(X1)|~r2_hidden(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26]), c_0_27])).
fof(c_0_31, plain, ![X66, X67, X68, X69]:(((r1_orders_2(X69,X68,X67)|~r1_lattice3(X69,k2_tarski(X67,X66),X68)|~epred1_4(X69,X68,X67,X66))&(r1_orders_2(X69,X68,X66)|~r1_lattice3(X69,k2_tarski(X67,X66),X68)|~epred1_4(X69,X68,X67,X66)))&((~r1_orders_2(X69,X68,X67)|~r1_orders_2(X69,X68,X66)|r1_lattice3(X69,k2_tarski(X67,X66),X68)|~epred1_4(X69,X68,X67,X66))&(((r1_orders_2(X69,X67,X68)|~r2_lattice3(X69,k2_tarski(X67,X66),X68)|~epred1_4(X69,X68,X67,X66))&(r1_orders_2(X69,X66,X68)|~r2_lattice3(X69,k2_tarski(X67,X66),X68)|~epred1_4(X69,X68,X67,X66)))&(~r1_orders_2(X69,X67,X68)|~r1_orders_2(X69,X66,X68)|r2_lattice3(X69,k2_tarski(X67,X66),X68)|~epred1_4(X69,X68,X67,X66))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_28])])])).
fof(c_0_32, plain, ![X22, X23]:k1_enumset1(X22,X22,X23)=k2_tarski(X22,X23), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_33, plain, ![X24, X25, X26]:k2_enumset1(X24,X24,X25,X26)=k1_enumset1(X24,X25,X26), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_34, plain, ![X27, X28, X29, X30]:k3_enumset1(X27,X27,X28,X29,X30)=k2_enumset1(X27,X28,X29,X30), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_35, plain, ![X31, X32, X33, X34, X35]:k4_enumset1(X31,X31,X32,X33,X34,X35)=k3_enumset1(X31,X32,X33,X34,X35), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_36, plain, ![X36, X37, X38, X39, X40, X41]:k5_enumset1(X36,X36,X37,X38,X39,X40,X41)=k4_enumset1(X36,X37,X38,X39,X40,X41), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_37, plain, ![X42, X43, X44, X45, X46, X47, X48]:k6_enumset1(X42,X42,X43,X44,X45,X46,X47,X48)=k5_enumset1(X42,X43,X44,X45,X46,X47,X48), inference(variable_rename,[status(thm)],[t75_enumset1])).
cnf(c_0_38, plain, (r1_lattice3(X1,X2,esk3_1(X1))|~v1_yellow_0(X1)|~l1_orders_2(X1)|~r2_hidden(esk2_3(X1,X2,esk3_1(X1)),u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_27])).
cnf(c_0_39, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_40, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,k3_yellow_0(X1),X2)))), inference(assume_negation,[status(cth)],[t44_yellow_0])).
cnf(c_0_41, plain, (r1_orders_2(X1,X2,X3)|~r1_lattice3(X1,k2_tarski(X4,X3),X2)|~epred1_4(X1,X2,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_42, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_44, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_46, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_47, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_48, plain, (r1_lattice3(X1,X2,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~v1_yellow_0(X1)|~l1_orders_2(X1)|~m1_subset_1(esk2_3(X1,X2,esk3_1(X1)),u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_49, plain, (m1_subset_1(esk2_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_50, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>epred1_4(X1,X2,X3,X4))))), inference(apply_def,[status(thm)],[t8_yellow_0, c_0_23])).
fof(c_0_51, negated_conjecture, ((~v2_struct_0(esk4_0)&(v5_orders_2(esk4_0)&(v1_yellow_0(esk4_0)&l1_orders_2(esk4_0))))&(m1_subset_1(esk5_0,u1_struct_0(esk4_0))&~r1_orders_2(esk4_0,k3_yellow_0(esk4_0),esk5_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_40])])])])).
fof(c_0_52, plain, ![X55]:(v2_struct_0(X55)|~l1_struct_0(X55)|~v1_xboole_0(u1_struct_0(X55))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_53, plain, (r1_orders_2(X1,X2,X3)|~epred1_4(X1,X2,X4,X3)|~r1_lattice3(X1,k6_enumset1(X4,X4,X4,X4,X4,X4,X4,X3),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_42]), c_0_43]), c_0_44]), c_0_45]), c_0_46]), c_0_47])).
cnf(c_0_54, plain, (r1_lattice3(X1,X2,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_27])).
fof(c_0_55, plain, ![X51, X52, X53, X54]:(~l1_orders_2(X51)|(~m1_subset_1(X52,u1_struct_0(X51))|(~m1_subset_1(X53,u1_struct_0(X51))|(~m1_subset_1(X54,u1_struct_0(X51))|epred1_4(X51,X52,X53,X54))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_50])])])).
cnf(c_0_56, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_57, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_52])).
fof(c_0_58, plain, ![X56]:(~l1_orders_2(X56)|l1_struct_0(X56)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_59, plain, (r1_orders_2(X1,esk3_1(X1),X2)|v1_xboole_0(u1_struct_0(X1))|~epred1_4(X1,esk3_1(X1),X3,X2)|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_60, plain, (epred1_4(X1,X2,X3,X4)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_61, negated_conjecture, (~l1_struct_0(esk4_0)|~v1_xboole_0(u1_struct_0(esk4_0))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_62, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_63, negated_conjecture, (l1_orders_2(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_51])).
fof(c_0_64, plain, ![X57, X58, X59, X60, X61]:(((r2_lattice3(X57,X59,X58)|(X58!=k1_yellow_0(X57,X59)|~r1_yellow_0(X57,X59))|~m1_subset_1(X58,u1_struct_0(X57))|(~v5_orders_2(X57)|~l1_orders_2(X57)))&(~m1_subset_1(X60,u1_struct_0(X57))|(~r2_lattice3(X57,X59,X60)|r1_orders_2(X57,X58,X60))|(X58!=k1_yellow_0(X57,X59)|~r1_yellow_0(X57,X59))|~m1_subset_1(X58,u1_struct_0(X57))|(~v5_orders_2(X57)|~l1_orders_2(X57))))&(((X58=k1_yellow_0(X57,X61)|(m1_subset_1(esk6_3(X57,X58,X61),u1_struct_0(X57))|~r2_lattice3(X57,X61,X58))|~m1_subset_1(X58,u1_struct_0(X57))|(~v5_orders_2(X57)|~l1_orders_2(X57)))&(r1_yellow_0(X57,X61)|(m1_subset_1(esk6_3(X57,X58,X61),u1_struct_0(X57))|~r2_lattice3(X57,X61,X58))|~m1_subset_1(X58,u1_struct_0(X57))|(~v5_orders_2(X57)|~l1_orders_2(X57))))&(((X58=k1_yellow_0(X57,X61)|(r2_lattice3(X57,X61,esk6_3(X57,X58,X61))|~r2_lattice3(X57,X61,X58))|~m1_subset_1(X58,u1_struct_0(X57))|(~v5_orders_2(X57)|~l1_orders_2(X57)))&(r1_yellow_0(X57,X61)|(r2_lattice3(X57,X61,esk6_3(X57,X58,X61))|~r2_lattice3(X57,X61,X58))|~m1_subset_1(X58,u1_struct_0(X57))|(~v5_orders_2(X57)|~l1_orders_2(X57))))&((X58=k1_yellow_0(X57,X61)|(~r1_orders_2(X57,X58,esk6_3(X57,X58,X61))|~r2_lattice3(X57,X61,X58))|~m1_subset_1(X58,u1_struct_0(X57))|(~v5_orders_2(X57)|~l1_orders_2(X57)))&(r1_yellow_0(X57,X61)|(~r1_orders_2(X57,X58,esk6_3(X57,X58,X61))|~r2_lattice3(X57,X61,X58))|~m1_subset_1(X58,u1_struct_0(X57))|(~v5_orders_2(X57)|~l1_orders_2(X57))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_yellow_0])])])])])])).
cnf(c_0_65, plain, (r1_orders_2(X1,esk3_1(X1),X2)|v1_xboole_0(u1_struct_0(X1))|~v1_yellow_0(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_27])).
cnf(c_0_66, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_67, negated_conjecture, (v1_yellow_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_68, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63])])).
cnf(c_0_69, plain, (r1_yellow_0(X1,X2)|~r1_orders_2(X1,X3,esk6_3(X1,X3,X2))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_70, negated_conjecture, (r1_orders_2(esk4_0,esk3_1(esk4_0),X1)|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67]), c_0_63])]), c_0_68])).
cnf(c_0_71, negated_conjecture, (v5_orders_2(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_72, plain, (X1=k1_yellow_0(X2,X3)|~r1_orders_2(X2,X1,esk6_3(X2,X1,X3))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_64])).
fof(c_0_73, plain, ![X63]:(~l1_orders_2(X63)|k3_yellow_0(X63)=k1_yellow_0(X63,k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_yellow_0])])).
cnf(c_0_74, negated_conjecture, (r1_yellow_0(esk4_0,X1)|~r2_lattice3(esk4_0,X1,esk3_1(esk4_0))|~m1_subset_1(esk6_3(esk4_0,esk3_1(esk4_0),X1),u1_struct_0(esk4_0))|~m1_subset_1(esk3_1(esk4_0),u1_struct_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_71]), c_0_63])])).
cnf(c_0_75, plain, (r1_yellow_0(X1,X2)|m1_subset_1(esk6_3(X1,X3,X2),u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
fof(c_0_76, plain, ![X64, X65]:((r2_lattice3(X64,k1_xboole_0,X65)|~m1_subset_1(X65,u1_struct_0(X64))|~l1_orders_2(X64))&(r1_lattice3(X64,k1_xboole_0,X65)|~m1_subset_1(X65,u1_struct_0(X64))|~l1_orders_2(X64))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_yellow_0])])])])).
cnf(c_0_77, negated_conjecture, (esk3_1(esk4_0)=k1_yellow_0(esk4_0,X1)|~r2_lattice3(esk4_0,X1,esk3_1(esk4_0))|~m1_subset_1(esk6_3(esk4_0,esk3_1(esk4_0),X1),u1_struct_0(esk4_0))|~m1_subset_1(esk3_1(esk4_0),u1_struct_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_70]), c_0_71]), c_0_63])])).
cnf(c_0_78, plain, (X1=k1_yellow_0(X2,X3)|m1_subset_1(esk6_3(X2,X1,X3),u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_79, negated_conjecture, (~r1_orders_2(esk4_0,k3_yellow_0(esk4_0),esk5_0)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_80, plain, (k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_81, plain, (r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~r1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_82, negated_conjecture, (r1_yellow_0(esk4_0,X1)|~r2_lattice3(esk4_0,X1,esk3_1(esk4_0))|~m1_subset_1(esk3_1(esk4_0),u1_struct_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_71]), c_0_63])])).
cnf(c_0_83, plain, (r2_lattice3(X1,k1_xboole_0,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_84, negated_conjecture, (esk3_1(esk4_0)=k1_yellow_0(esk4_0,X1)|~r2_lattice3(esk4_0,X1,esk3_1(esk4_0))|~m1_subset_1(esk3_1(esk4_0),u1_struct_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_71]), c_0_63])])).
cnf(c_0_85, negated_conjecture, (~r1_orders_2(esk4_0,k1_yellow_0(esk4_0,k1_xboole_0),esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_63])])).
cnf(c_0_86, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),X3)|~r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(er,[status(thm)],[c_0_81])).
cnf(c_0_87, negated_conjecture, (r1_yellow_0(esk4_0,k1_xboole_0)|~m1_subset_1(esk3_1(esk4_0),u1_struct_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_63])])).
cnf(c_0_88, negated_conjecture, (esk3_1(esk4_0)=k1_yellow_0(esk4_0,k1_xboole_0)|~m1_subset_1(esk3_1(esk4_0),u1_struct_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_83]), c_0_63])])).
cnf(c_0_89, negated_conjecture, (~r1_yellow_0(esk4_0,k1_xboole_0)|~r2_lattice3(esk4_0,k1_xboole_0,esk5_0)|~m1_subset_1(k1_yellow_0(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_71]), c_0_63]), c_0_66])])).
cnf(c_0_90, negated_conjecture, (r1_yellow_0(esk4_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_27]), c_0_67]), c_0_63])])).
cnf(c_0_91, negated_conjecture, (esk3_1(esk4_0)=k1_yellow_0(esk4_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_27]), c_0_67]), c_0_63])])).
cnf(c_0_92, negated_conjecture, (~r2_lattice3(esk4_0,k1_xboole_0,esk5_0)|~m1_subset_1(k1_yellow_0(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_90])])).
cnf(c_0_93, negated_conjecture, (m1_subset_1(k1_yellow_0(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_91]), c_0_67]), c_0_63])])).
cnf(c_0_94, negated_conjecture, (~r2_lattice3(esk4_0,k1_xboole_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_93])])).
cnf(c_0_95, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_83]), c_0_63]), c_0_66])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 96
# Proof object clause steps            : 57
# Proof object formula steps           : 39
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 57
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 45
# Removed in clause preprocessing      : 6
# Initial clauses in saturation        : 39
# Processed clauses                    : 140
# ...of these trivial                  : 0
# ...subsumed                          : 13
# ...remaining for further processing  : 127
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 10
# Generated clauses                    : 124
# ...of the previous two non-trivial   : 113
# Contextual simplify-reflections      : 6
# Paramodulations                      : 122
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 71
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 58
# Current number of unprocessed clauses: 47
# ...number of literals in the above   : 255
# Current number of archived formulas  : 0
# Current number of archived clauses   : 60
# Clause-clause subsumption calls (NU) : 1185
# Rec. Clause-clause subsumption calls : 279
# Non-unit clause-clause subsumptions  : 23
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6273

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.000 s
# Total time               : 0.015 s
# Maximum resident set size: 28696 pages
