# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t178_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>r1_tarski(k10_relat_1(X3,X1),k10_relat_1(X3,X2)))), file('data/MPTP/proofs/25002676c78346c9972e83600c1a00f1.E_input', t178_relat_1)).
fof(t169_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k10_relat_1(X1,k2_relat_1(X1))=k1_relat_1(X1)), file('data/MPTP/proofs/25002676c78346c9972e83600c1a00f1.E_input', t169_relat_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/25002676c78346c9972e83600c1a00f1.E_input', d3_tarski)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/25002676c78346c9972e83600c1a00f1.E_input', d10_xboole_0)).
fof(t20_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)=>(r2_hidden(X1,k1_relat_1(X3))&r2_hidden(X2,k2_relat_1(X3))))), file('data/MPTP/proofs/25002676c78346c9972e83600c1a00f1.E_input', t20_relat_1)).
fof(t166_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k10_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k2_relat_1(X3))&(r2_hidden(k4_tarski(X1,X4),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/25002676c78346c9972e83600c1a00f1.E_input', t166_relat_1)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/25002676c78346c9972e83600c1a00f1.E_input', t1_xboole_1)).
fof(t181_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>k10_relat_1(k5_relat_1(X2,X3),X1)=k10_relat_1(X2,k10_relat_1(X3,X1)))), file('data/MPTP/proofs/25002676c78346c9972e83600c1a00f1.E_input', t181_relat_1)).
fof(dt_k5_relat_1, axiom, ![X1, X2]:((v1_relat_1(X1)&v1_relat_1(X2))=>v1_relat_1(k5_relat_1(X1,X2))), file('data/MPTP/proofs/25002676c78346c9972e83600c1a00f1.E_input', dt_k5_relat_1)).
fof(t182_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>k1_relat_1(k5_relat_1(X1,X2))=k10_relat_1(X1,k1_relat_1(X2)))), file('data/MPTP/proofs/25002676c78346c9972e83600c1a00f1.E_input', t182_relat_1)).
fof(c_0_10, plain, ![X5, X6, X7]:(~v1_relat_1(X7)|(~r1_tarski(X5,X6)|r1_tarski(k10_relat_1(X7,X5),k10_relat_1(X7,X6)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t178_relat_1])])).
fof(c_0_11, plain, ![X8]:(~v1_relat_1(X8)|k10_relat_1(X8,k2_relat_1(X8))=k1_relat_1(X8)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t169_relat_1])])).
fof(c_0_12, plain, ![X9, X10, X11, X12, X13]:((~r1_tarski(X9,X10)|(~r2_hidden(X11,X9)|r2_hidden(X11,X10)))&((r2_hidden(esk1_2(X12,X13),X12)|r1_tarski(X12,X13))&(~r2_hidden(esk1_2(X12,X13),X13)|r1_tarski(X12,X13)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_13, plain, (r1_tarski(k10_relat_1(X1,X2),k10_relat_1(X1,X3))|~v1_relat_1(X1)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (k10_relat_1(X1,k2_relat_1(X1))=k1_relat_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_15, plain, ![X15, X16]:(((r1_tarski(X15,X16)|X15!=X16)&(r1_tarski(X16,X15)|X15!=X16))&(~r1_tarski(X15,X16)|~r1_tarski(X16,X15)|X15=X16)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_16, plain, ![X17, X18, X19]:((r2_hidden(X17,k1_relat_1(X19))|~r2_hidden(k4_tarski(X17,X18),X19)|~v1_relat_1(X19))&(r2_hidden(X18,k2_relat_1(X19))|~r2_hidden(k4_tarski(X17,X18),X19)|~v1_relat_1(X19))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t20_relat_1])])])).
fof(c_0_17, plain, ![X20, X21, X22, X24]:(((r2_hidden(esk2_3(X20,X21,X22),k2_relat_1(X22))|~r2_hidden(X20,k10_relat_1(X22,X21))|~v1_relat_1(X22))&((r2_hidden(k4_tarski(X20,esk2_3(X20,X21,X22)),X22)|~r2_hidden(X20,k10_relat_1(X22,X21))|~v1_relat_1(X22))&(r2_hidden(esk2_3(X20,X21,X22),X21)|~r2_hidden(X20,k10_relat_1(X22,X21))|~v1_relat_1(X22))))&(~r2_hidden(X24,k2_relat_1(X22))|(~r2_hidden(k4_tarski(X20,X24),X22)|~r2_hidden(X24,X21))|r2_hidden(X20,k10_relat_1(X22,X21))|~v1_relat_1(X22))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t166_relat_1])])])])])).
cnf(c_0_18, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, plain, (r1_tarski(k1_relat_1(X1),k10_relat_1(X1,X2))|~r1_tarski(k2_relat_1(X1),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_20, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (r2_hidden(k4_tarski(X1,esk2_3(X1,X2,X3)),X3)|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (r2_hidden(X1,k10_relat_1(X2,X3))|~r2_hidden(X1,k1_relat_1(X2))|~r1_tarski(k2_relat_1(X2),X3)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_24, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_20])).
fof(c_0_25, plain, ![X25, X26, X27]:(~r1_tarski(X25,X26)|~r1_tarski(X26,X27)|r1_tarski(X25,X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
cnf(c_0_26, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_27, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(X1,k10_relat_1(X2,X3))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28, plain, (r2_hidden(esk1_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_29, plain, (r1_tarski(X1,X2)|~r2_hidden(esk1_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_30, plain, (r2_hidden(X1,k10_relat_1(X2,k2_relat_1(X2)))|~r2_hidden(X1,k1_relat_1(X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_31, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, plain, (k10_relat_1(X1,X2)=k10_relat_1(X1,X3)|~r1_tarski(k10_relat_1(X1,X2),k10_relat_1(X1,X3))|~r1_tarski(X3,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_26, c_0_13])).
cnf(c_0_33, plain, (r2_hidden(esk1_2(k10_relat_1(X1,X2),X3),k1_relat_1(X1))|r1_tarski(k10_relat_1(X1,X2),X3)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34, plain, (r1_tarski(X1,k10_relat_1(X2,k2_relat_1(X2)))|~r2_hidden(esk1_2(X1,k10_relat_1(X2,k2_relat_1(X2))),k1_relat_1(X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_35, plain, (r1_tarski(X1,k10_relat_1(X2,X3))|~r1_tarski(X1,k10_relat_1(X2,X4))|~r1_tarski(X4,X3)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_31, c_0_13])).
cnf(c_0_36, plain, (k10_relat_1(X1,X2)=k10_relat_1(X1,X3)|~r1_tarski(X3,X2)|~r1_tarski(X2,X3)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_32, c_0_13])).
cnf(c_0_37, plain, (r1_tarski(k10_relat_1(X1,X2),k1_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_29, c_0_33])).
cnf(c_0_38, plain, (r1_tarski(k1_relat_1(X1),k10_relat_1(X1,k2_relat_1(X1)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_34, c_0_28])).
cnf(c_0_39, plain, (r1_tarski(k10_relat_1(X1,X2),k10_relat_1(X1,X3))|~r1_tarski(X4,X3)|~r1_tarski(X2,X4)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_35, c_0_13])).
fof(c_0_40, plain, ![X28, X29, X30]:(~v1_relat_1(X29)|(~v1_relat_1(X30)|k10_relat_1(k5_relat_1(X29,X30),X28)=k10_relat_1(X29,k10_relat_1(X30,X28)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t181_relat_1])])])).
fof(c_0_41, plain, ![X31, X32]:(~v1_relat_1(X31)|~v1_relat_1(X32)|v1_relat_1(k5_relat_1(X31,X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_relat_1])])).
cnf(c_0_42, plain, (k10_relat_1(X1,k1_relat_1(X2))=k10_relat_1(X1,k10_relat_1(X2,X3))|~r1_tarski(k1_relat_1(X2),k10_relat_1(X2,X3))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_43, plain, (r1_tarski(X1,k10_relat_1(X2,k2_relat_1(X2)))|~r1_tarski(X1,k1_relat_1(X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_31, c_0_38])).
cnf(c_0_44, plain, (r1_tarski(k10_relat_1(X1,X2),k10_relat_1(X1,k1_relat_1(X3)))|~r1_tarski(X2,k10_relat_1(X3,X4))|~v1_relat_1(X1)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_39, c_0_37])).
cnf(c_0_45, plain, (r1_tarski(k10_relat_1(X1,X2),k10_relat_1(X1,k2_relat_1(X1)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_34, c_0_33])).
cnf(c_0_46, plain, (k10_relat_1(k5_relat_1(X1,X2),X3)=k10_relat_1(X1,k10_relat_1(X2,X3))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_47, plain, (v1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_48, plain, (k10_relat_1(X1,k10_relat_1(X2,k2_relat_1(X2)))=k10_relat_1(X1,k1_relat_1(X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_24])])).
cnf(c_0_49, plain, (r1_tarski(k10_relat_1(X1,k10_relat_1(X2,X3)),k10_relat_1(X1,k1_relat_1(X2)))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_50, plain, (k10_relat_1(X1,k10_relat_1(X2,k2_relat_1(k5_relat_1(X1,X2))))=k1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_46]), c_0_47])).
cnf(c_0_51, plain, (r1_tarski(X1,k1_relat_1(X2))|~r1_tarski(X1,k10_relat_1(X2,X3))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_31, c_0_37])).
cnf(c_0_52, plain, (r1_tarski(k10_relat_1(X1,k1_relat_1(X2)),k10_relat_1(X1,X3))|~r1_tarski(k10_relat_1(X2,k2_relat_1(X2)),X3)|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_13, c_0_48])).
fof(c_0_53, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>k1_relat_1(k5_relat_1(X1,X2))=k10_relat_1(X1,k1_relat_1(X2))))), inference(assume_negation,[status(cth)],[t182_relat_1])).
cnf(c_0_54, plain, (r1_tarski(k1_relat_1(k5_relat_1(X1,X2)),k10_relat_1(X1,k1_relat_1(X2)))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_55, plain, (r1_tarski(X1,k1_relat_1(k5_relat_1(X2,X3)))|~r1_tarski(X1,k10_relat_1(X2,k10_relat_1(X3,X4)))|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_46]), c_0_47])).
cnf(c_0_56, plain, (r1_tarski(k10_relat_1(X1,k1_relat_1(X2)),k10_relat_1(X1,k10_relat_1(X2,k2_relat_1(X2))))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_52, c_0_24])).
fof(c_0_57, negated_conjecture, (v1_relat_1(esk3_0)&(v1_relat_1(esk4_0)&k1_relat_1(k5_relat_1(esk3_0,esk4_0))!=k10_relat_1(esk3_0,k1_relat_1(esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_53])])])).
cnf(c_0_58, plain, (k1_relat_1(k5_relat_1(X1,X2))=k10_relat_1(X1,k1_relat_1(X2))|~r1_tarski(k10_relat_1(X1,k1_relat_1(X2)),k1_relat_1(k5_relat_1(X1,X2)))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_26, c_0_54])).
cnf(c_0_59, plain, (r1_tarski(k10_relat_1(X1,k1_relat_1(X2)),k1_relat_1(k5_relat_1(X1,X2)))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_60, negated_conjecture, (k1_relat_1(k5_relat_1(esk3_0,esk4_0))!=k10_relat_1(esk3_0,k1_relat_1(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_61, plain, (k1_relat_1(k5_relat_1(X1,X2))=k10_relat_1(X1,k1_relat_1(X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_62, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_63, negated_conjecture, (v1_relat_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_64, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62]), c_0_63])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 65
# Proof object clause steps            : 44
# Proof object formula steps           : 21
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 28
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 20
# Processed clauses                    : 541
# ...of these trivial                  : 0
# ...subsumed                          : 324
# ...remaining for further processing  : 217
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 14
# Backward-rewritten                   : 0
# Generated clauses                    : 2496
# ...of the previous two non-trivial   : 2384
# Contextual simplify-reflections      : 18
# Paramodulations                      : 2494
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 182
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 178
# Current number of unprocessed clauses: 1816
# ...number of literals in the above   : 10782
# Current number of archived formulas  : 0
# Current number of archived clauses   : 33
# Clause-clause subsumption calls (NU) : 6527
# Rec. Clause-clause subsumption calls : 2174
# Non-unit clause-clause subsumptions  : 356
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 49813

# -------------------------------------------------
# User time                : 0.039 s
# System time              : 0.000 s
# Total time               : 0.039 s
# Maximum resident set size: 28776 pages
