# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t2_filter_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_filter_0(X1,X2)=k2_filter_0(X1,X3)=>X2=X3)))), file('data/MPTP/proofs/2540e359eea3414e831613b12342fe66.E_input', t2_filter_1)).
fof(t18_filter_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X2,k2_filter_0(X1,X3))<=>r3_lattices(X1,X3,X2))))), file('data/MPTP/proofs/2540e359eea3414e831613b12342fe66.E_input', t18_filter_0)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/2540e359eea3414e831613b12342fe66.E_input', cc1_lattices)).
fof(reflexivity_r3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>r3_lattices(X1,X2,X2)), file('data/MPTP/proofs/2540e359eea3414e831613b12342fe66.E_input', reflexivity_r3_lattices)).
fof(dt_l3_lattices, axiom, ![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1))), file('data/MPTP/proofs/2540e359eea3414e831613b12342fe66.E_input', dt_l3_lattices)).
fof(t26_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&(v4_lattices(X1)&l2_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattices(X1,X2,X3)&r1_lattices(X1,X3,X2))=>X2=X3)))), file('data/MPTP/proofs/2540e359eea3414e831613b12342fe66.E_input', t26_lattices)).
fof(redefinition_r3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_lattices(X1,X2,X3)<=>r1_lattices(X1,X2,X3))), file('data/MPTP/proofs/2540e359eea3414e831613b12342fe66.E_input', redefinition_r3_lattices)).
fof(c_0_7, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_filter_0(X1,X2)=k2_filter_0(X1,X3)=>X2=X3))))), inference(assume_negation,[status(cth)],[t2_filter_1])).
fof(c_0_8, plain, ![X7, X8, X9]:((~r2_hidden(X8,k2_filter_0(X7,X9))|r3_lattices(X7,X9,X8)|~m1_subset_1(X9,u1_struct_0(X7))|~m1_subset_1(X8,u1_struct_0(X7))|(v2_struct_0(X7)|(~v10_lattices(X7)|~l3_lattices(X7))))&(~r3_lattices(X7,X9,X8)|r2_hidden(X8,k2_filter_0(X7,X9))|~m1_subset_1(X9,u1_struct_0(X7))|~m1_subset_1(X8,u1_struct_0(X7))|(v2_struct_0(X7)|(~v10_lattices(X7)|~l3_lattices(X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t18_filter_0])])])])])).
fof(c_0_9, negated_conjecture, ((~v2_struct_0(esk1_0)&(v10_lattices(esk1_0)&l3_lattices(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(k2_filter_0(esk1_0,esk2_0)=k2_filter_0(esk1_0,esk3_0)&esk2_0!=esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_7])])])])).
fof(c_0_10, plain, ![X16]:((~v2_struct_0(X16)|(v2_struct_0(X16)|~v10_lattices(X16))|~l3_lattices(X16))&((v4_lattices(X16)|(v2_struct_0(X16)|~v10_lattices(X16))|~l3_lattices(X16))&((v5_lattices(X16)|(v2_struct_0(X16)|~v10_lattices(X16))|~l3_lattices(X16))&((v6_lattices(X16)|(v2_struct_0(X16)|~v10_lattices(X16))|~l3_lattices(X16))&((v7_lattices(X16)|(v2_struct_0(X16)|~v10_lattices(X16))|~l3_lattices(X16))&((v8_lattices(X16)|(v2_struct_0(X16)|~v10_lattices(X16))|~l3_lattices(X16))&(v9_lattices(X16)|(v2_struct_0(X16)|~v10_lattices(X16))|~l3_lattices(X16)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
cnf(c_0_11, plain, (r2_hidden(X3,k2_filter_0(X1,X2))|v2_struct_0(X1)|~r3_lattices(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, negated_conjecture, (l3_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, negated_conjecture, (v10_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_16, plain, ![X18, X19, X20]:(v2_struct_0(X18)|(~v6_lattices(X18)|(~v8_lattices(X18)|(~v9_lattices(X18)|~l3_lattices(X18))))|(~m1_subset_1(X19,u1_struct_0(X18))|~m1_subset_1(X20,u1_struct_0(X18)))|r3_lattices(X18,X19,X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r3_lattices])])])).
cnf(c_0_17, plain, (v9_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, plain, (v8_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19, plain, (v6_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_20, plain, ![X17]:((l1_lattices(X17)|~l3_lattices(X17))&(l2_lattices(X17)|~l3_lattices(X17))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
cnf(c_0_21, negated_conjecture, (r2_hidden(esk3_0,k2_filter_0(esk1_0,X1))|~r3_lattices(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13]), c_0_14])]), c_0_15])).
cnf(c_0_22, negated_conjecture, (k2_filter_0(esk1_0,esk2_0)=k2_filter_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_23, plain, (v2_struct_0(X1)|r3_lattices(X1,X2,X2)|~v6_lattices(X1)|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, negated_conjecture, (v9_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_14]), c_0_13])]), c_0_15])).
cnf(c_0_25, negated_conjecture, (v8_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_14]), c_0_13])]), c_0_15])).
cnf(c_0_26, negated_conjecture, (v6_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_14]), c_0_13])]), c_0_15])).
cnf(c_0_27, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_28, plain, ![X10, X11, X12]:(v2_struct_0(X10)|(~v4_lattices(X10)|~l2_lattices(X10))|(~m1_subset_1(X11,u1_struct_0(X10))|(~m1_subset_1(X12,u1_struct_0(X10))|(~r1_lattices(X10,X11,X12)|~r1_lattices(X10,X12,X11)|X11=X12)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t26_lattices])])])])).
cnf(c_0_29, plain, (l2_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30, plain, (v4_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_31, plain, ![X13, X14, X15]:((~r3_lattices(X13,X14,X15)|r1_lattices(X13,X14,X15)|(v2_struct_0(X13)|(~v6_lattices(X13)|(~v8_lattices(X13)|(~v9_lattices(X13)|~l3_lattices(X13))))|(~m1_subset_1(X14,u1_struct_0(X13))|~m1_subset_1(X15,u1_struct_0(X13)))))&(~r1_lattices(X13,X14,X15)|r3_lattices(X13,X14,X15)|(v2_struct_0(X13)|(~v6_lattices(X13)|(~v8_lattices(X13)|(~v9_lattices(X13)|~l3_lattices(X13))))|(~m1_subset_1(X14,u1_struct_0(X13))|~m1_subset_1(X15,u1_struct_0(X13)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_lattices])])])])).
cnf(c_0_32, plain, (r3_lattices(X2,X3,X1)|v2_struct_0(X2)|~r2_hidden(X1,k2_filter_0(X2,X3))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~v10_lattices(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_33, negated_conjecture, (r2_hidden(esk3_0,k2_filter_0(esk1_0,esk2_0))|~r3_lattices(esk1_0,esk3_0,esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_12]), c_0_22])).
cnf(c_0_34, negated_conjecture, (r3_lattices(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_12]), c_0_13])]), c_0_15]), c_0_24]), c_0_25]), c_0_26])])).
cnf(c_0_35, negated_conjecture, (r2_hidden(esk2_0,k2_filter_0(esk1_0,X1))|~r3_lattices(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_27]), c_0_13]), c_0_14])]), c_0_15])).
cnf(c_0_36, plain, (v2_struct_0(X1)|X2=X3|~v4_lattices(X1)|~l2_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattices(X1,X2,X3)|~r1_lattices(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, negated_conjecture, (l2_lattices(esk1_0)), inference(spm,[status(thm)],[c_0_29, c_0_13])).
cnf(c_0_38, negated_conjecture, (v4_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_14]), c_0_13])]), c_0_15])).
cnf(c_0_39, plain, (r1_lattices(X1,X2,X3)|v2_struct_0(X1)|~r3_lattices(X1,X2,X3)|~v6_lattices(X1)|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_40, negated_conjecture, (r3_lattices(esk1_0,esk2_0,esk3_0)|~r3_lattices(esk1_0,esk3_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_27]), c_0_12]), c_0_13]), c_0_14])]), c_0_15])).
cnf(c_0_41, negated_conjecture, (r3_lattices(esk1_0,esk3_0,esk3_0)), inference(spm,[status(thm)],[c_0_34, c_0_12])).
cnf(c_0_42, negated_conjecture, (r3_lattices(esk1_0,esk3_0,X1)|~r2_hidden(X1,k2_filter_0(esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_22]), c_0_12]), c_0_13]), c_0_14])]), c_0_15])).
cnf(c_0_43, negated_conjecture, (r2_hidden(esk2_0,k2_filter_0(esk1_0,esk2_0))|~r3_lattices(esk1_0,esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_35, c_0_27])).
cnf(c_0_44, negated_conjecture, (r3_lattices(esk1_0,esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_34, c_0_27])).
cnf(c_0_45, negated_conjecture, (X1=esk3_0|~r1_lattices(esk1_0,esk3_0,X1)|~r1_lattices(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_12]), c_0_15]), c_0_37]), c_0_38])])).
cnf(c_0_46, negated_conjecture, (esk2_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_47, negated_conjecture, (r1_lattices(esk1_0,X1,esk3_0)|~r3_lattices(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_12]), c_0_13])]), c_0_15]), c_0_24]), c_0_25]), c_0_26])])).
cnf(c_0_48, negated_conjecture, (r3_lattices(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_41])])).
cnf(c_0_49, negated_conjecture, (r3_lattices(esk1_0,esk3_0,esk2_0)|~r2_hidden(esk2_0,k2_filter_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_42, c_0_27])).
cnf(c_0_50, negated_conjecture, (r2_hidden(esk2_0,k2_filter_0(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_44])])).
cnf(c_0_51, negated_conjecture, (~r1_lattices(esk1_0,esk3_0,esk2_0)|~r1_lattices(esk1_0,esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_27]), c_0_46])).
cnf(c_0_52, negated_conjecture, (r1_lattices(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_27]), c_0_48])])).
cnf(c_0_53, negated_conjecture, (r1_lattices(esk1_0,X1,esk2_0)|~r3_lattices(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_27]), c_0_13])]), c_0_15]), c_0_24]), c_0_25]), c_0_26])])).
cnf(c_0_54, negated_conjecture, (r3_lattices(esk1_0,esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_50])])).
cnf(c_0_55, negated_conjecture, (~r1_lattices(esk1_0,esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_52])])).
cnf(c_0_56, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_12]), c_0_54])]), c_0_55]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 57
# Proof object clause steps            : 42
# Proof object formula steps           : 15
# Proof object conjectures             : 35
# Proof object clause conjectures      : 32
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 71
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 21
# Processed clauses                    : 82
# ...of these trivial                  : 1
# ...subsumed                          : 2
# ...remaining for further processing  : 79
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 8
# Generated clauses                    : 48
# ...of the previous two non-trivial   : 41
# Contextual simplify-reflections      : 0
# Paramodulations                      : 48
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 48
#    Positive orientable unit clauses  : 21
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 24
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 31
# Clause-clause subsumption calls (NU) : 753
# Rec. Clause-clause subsumption calls : 138
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 101
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 14
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2932

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.003 s
# Total time               : 0.010 s
# Maximum resident set size: 29024 pages
