# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/2566d725c3a04e7ab76eae81217dfd5c.E_input', d10_xboole_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/2566d725c3a04e7ab76eae81217dfd5c.E_input', t3_subset)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/2566d725c3a04e7ab76eae81217dfd5c.E_input', t5_subset)).
fof(dt_u1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/2566d725c3a04e7ab76eae81217dfd5c.E_input', dt_u1_pre_topc)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/2566d725c3a04e7ab76eae81217dfd5c.E_input', d2_subset_1)).
fof(t24_yellow_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1)), file('data/MPTP/proofs/2566d725c3a04e7ab76eae81217dfd5c.E_input', t24_yellow_1)).
fof(t94_zfmisc_1, axiom, ![X1, X2]:(![X3]:(r2_hidden(X3,X1)=>r1_tarski(X3,X2))=>r1_tarski(k3_tarski(X1),X2)), file('data/MPTP/proofs/2566d725c3a04e7ab76eae81217dfd5c.E_input', t94_zfmisc_1)).
fof(d1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1))))))))), file('data/MPTP/proofs/2566d725c3a04e7ab76eae81217dfd5c.E_input', d1_pre_topc)).
fof(l49_zfmisc_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>r1_tarski(X1,k3_tarski(X2))), file('data/MPTP/proofs/2566d725c3a04e7ab76eae81217dfd5c.E_input', l49_zfmisc_1)).
fof(t95_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>r1_tarski(k3_tarski(X1),k3_tarski(X2))), file('data/MPTP/proofs/2566d725c3a04e7ab76eae81217dfd5c.E_input', t95_zfmisc_1)).
fof(t14_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>(r2_hidden(k3_tarski(X1),X1)=>k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1))), file('data/MPTP/proofs/2566d725c3a04e7ab76eae81217dfd5c.E_input', t14_yellow_1)).
fof(c_0_11, plain, ![X6, X7]:(((r1_tarski(X6,X7)|X6!=X7)&(r1_tarski(X7,X6)|X6!=X7))&(~r1_tarski(X6,X7)|~r1_tarski(X7,X6)|X6=X7)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_12, plain, ![X28, X29]:((~m1_subset_1(X28,k1_zfmisc_1(X29))|r1_tarski(X28,X29))&(~r1_tarski(X28,X29)|m1_subset_1(X28,k1_zfmisc_1(X29)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_13, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_14, plain, ![X25, X26, X27]:(~r2_hidden(X25,X26)|(~m1_subset_1(X26,k1_zfmisc_1(X27))|~v1_xboole_0(X27))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
cnf(c_0_15, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_13])).
fof(c_0_17, plain, ![X5]:(~l1_pre_topc(X5)|m1_subset_1(u1_pre_topc(X5),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X5))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_pre_topc])])).
fof(c_0_18, plain, ![X10, X11]:(((~m1_subset_1(X11,X10)|r2_hidden(X11,X10)|v1_xboole_0(X10))&(~r2_hidden(X11,X10)|m1_subset_1(X11,X10)|v1_xboole_0(X10)))&((~m1_subset_1(X11,X10)|v1_xboole_0(X11)|~v1_xboole_0(X10))&(~v1_xboole_0(X11)|m1_subset_1(X11,X10)|~v1_xboole_0(X10)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_19, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1))), inference(assume_negation,[status(cth)],[t24_yellow_1])).
fof(c_0_20, plain, ![X22, X23]:((r2_hidden(esk5_2(X22,X23),X22)|r1_tarski(k3_tarski(X22),X23))&(~r1_tarski(esk5_2(X22,X23),X23)|r1_tarski(k3_tarski(X22),X23))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t94_zfmisc_1])])])])).
cnf(c_0_21, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_23, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_25, plain, ![X13, X14, X15, X16]:(((r2_hidden(u1_struct_0(X13),u1_pre_topc(X13))|~v2_pre_topc(X13)|~l1_pre_topc(X13))&((~m1_subset_1(X14,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X13))))|(~r1_tarski(X14,u1_pre_topc(X13))|r2_hidden(k5_setfam_1(u1_struct_0(X13),X14),u1_pre_topc(X13)))|~v2_pre_topc(X13)|~l1_pre_topc(X13))&(~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X13)))|(~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X13)))|(~r2_hidden(X15,u1_pre_topc(X13))|~r2_hidden(X16,u1_pre_topc(X13))|r2_hidden(k9_subset_1(u1_struct_0(X13),X15,X16),u1_pre_topc(X13))))|~v2_pre_topc(X13)|~l1_pre_topc(X13))))&(((m1_subset_1(esk3_1(X13),k1_zfmisc_1(u1_struct_0(X13)))|m1_subset_1(esk2_1(X13),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X13))))|~r2_hidden(u1_struct_0(X13),u1_pre_topc(X13))|v2_pre_topc(X13)|~l1_pre_topc(X13))&((m1_subset_1(esk4_1(X13),k1_zfmisc_1(u1_struct_0(X13)))|m1_subset_1(esk2_1(X13),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X13))))|~r2_hidden(u1_struct_0(X13),u1_pre_topc(X13))|v2_pre_topc(X13)|~l1_pre_topc(X13))&(((r2_hidden(esk3_1(X13),u1_pre_topc(X13))|m1_subset_1(esk2_1(X13),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X13))))|~r2_hidden(u1_struct_0(X13),u1_pre_topc(X13))|v2_pre_topc(X13)|~l1_pre_topc(X13))&(r2_hidden(esk4_1(X13),u1_pre_topc(X13))|m1_subset_1(esk2_1(X13),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X13))))|~r2_hidden(u1_struct_0(X13),u1_pre_topc(X13))|v2_pre_topc(X13)|~l1_pre_topc(X13)))&(~r2_hidden(k9_subset_1(u1_struct_0(X13),esk3_1(X13),esk4_1(X13)),u1_pre_topc(X13))|m1_subset_1(esk2_1(X13),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X13))))|~r2_hidden(u1_struct_0(X13),u1_pre_topc(X13))|v2_pre_topc(X13)|~l1_pre_topc(X13)))))&(((m1_subset_1(esk3_1(X13),k1_zfmisc_1(u1_struct_0(X13)))|r1_tarski(esk2_1(X13),u1_pre_topc(X13))|~r2_hidden(u1_struct_0(X13),u1_pre_topc(X13))|v2_pre_topc(X13)|~l1_pre_topc(X13))&((m1_subset_1(esk4_1(X13),k1_zfmisc_1(u1_struct_0(X13)))|r1_tarski(esk2_1(X13),u1_pre_topc(X13))|~r2_hidden(u1_struct_0(X13),u1_pre_topc(X13))|v2_pre_topc(X13)|~l1_pre_topc(X13))&(((r2_hidden(esk3_1(X13),u1_pre_topc(X13))|r1_tarski(esk2_1(X13),u1_pre_topc(X13))|~r2_hidden(u1_struct_0(X13),u1_pre_topc(X13))|v2_pre_topc(X13)|~l1_pre_topc(X13))&(r2_hidden(esk4_1(X13),u1_pre_topc(X13))|r1_tarski(esk2_1(X13),u1_pre_topc(X13))|~r2_hidden(u1_struct_0(X13),u1_pre_topc(X13))|v2_pre_topc(X13)|~l1_pre_topc(X13)))&(~r2_hidden(k9_subset_1(u1_struct_0(X13),esk3_1(X13),esk4_1(X13)),u1_pre_topc(X13))|r1_tarski(esk2_1(X13),u1_pre_topc(X13))|~r2_hidden(u1_struct_0(X13),u1_pre_topc(X13))|v2_pre_topc(X13)|~l1_pre_topc(X13)))))&((m1_subset_1(esk3_1(X13),k1_zfmisc_1(u1_struct_0(X13)))|~r2_hidden(k5_setfam_1(u1_struct_0(X13),esk2_1(X13)),u1_pre_topc(X13))|~r2_hidden(u1_struct_0(X13),u1_pre_topc(X13))|v2_pre_topc(X13)|~l1_pre_topc(X13))&((m1_subset_1(esk4_1(X13),k1_zfmisc_1(u1_struct_0(X13)))|~r2_hidden(k5_setfam_1(u1_struct_0(X13),esk2_1(X13)),u1_pre_topc(X13))|~r2_hidden(u1_struct_0(X13),u1_pre_topc(X13))|v2_pre_topc(X13)|~l1_pre_topc(X13))&(((r2_hidden(esk3_1(X13),u1_pre_topc(X13))|~r2_hidden(k5_setfam_1(u1_struct_0(X13),esk2_1(X13)),u1_pre_topc(X13))|~r2_hidden(u1_struct_0(X13),u1_pre_topc(X13))|v2_pre_topc(X13)|~l1_pre_topc(X13))&(r2_hidden(esk4_1(X13),u1_pre_topc(X13))|~r2_hidden(k5_setfam_1(u1_struct_0(X13),esk2_1(X13)),u1_pre_topc(X13))|~r2_hidden(u1_struct_0(X13),u1_pre_topc(X13))|v2_pre_topc(X13)|~l1_pre_topc(X13)))&(~r2_hidden(k9_subset_1(u1_struct_0(X13),esk3_1(X13),esk4_1(X13)),u1_pre_topc(X13))|~r2_hidden(k5_setfam_1(u1_struct_0(X13),esk2_1(X13)),u1_pre_topc(X13))|~r2_hidden(u1_struct_0(X13),u1_pre_topc(X13))|v2_pre_topc(X13)|~l1_pre_topc(X13)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])).
fof(c_0_26, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&k4_yellow_0(k2_yellow_1(u1_pre_topc(esk1_0)))!=u1_struct_0(esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_19])])])])).
cnf(c_0_27, plain, (r1_tarski(k3_tarski(X1),X2)|~r1_tarski(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_29, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_30, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_31, plain, (~v1_xboole_0(k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,u1_pre_topc(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_21, c_0_23])).
cnf(c_0_32, plain, (v1_xboole_0(k1_zfmisc_1(X1))|r2_hidden(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_24, c_0_22])).
cnf(c_0_33, plain, (r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_36, plain, ![X8, X9]:(~r2_hidden(X8,X9)|r1_tarski(X8,k3_tarski(X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l49_zfmisc_1])])).
cnf(c_0_37, plain, (r1_tarski(k3_tarski(X1),X2)|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_38, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_39, plain, (r2_hidden(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,u1_pre_topc(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_40, negated_conjecture, (r2_hidden(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])])).
cnf(c_0_41, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_42, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_43, plain, (r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44, plain, (r1_tarski(k3_tarski(X1),X2)|~r2_hidden(esk5_2(X1,X2),k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_45, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(k3_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_46, negated_conjecture, (r2_hidden(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_35])])).
cnf(c_0_47, plain, (v1_xboole_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_41, c_0_22])).
cnf(c_0_48, plain, (k3_tarski(X1)=X2|~r2_hidden(X2,X1)|~r1_tarski(k3_tarski(X1),X2)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_49, plain, (r1_tarski(k3_tarski(k1_zfmisc_1(X1)),X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_50, negated_conjecture, (~v1_xboole_0(k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_30, c_0_46])).
cnf(c_0_51, plain, (v1_xboole_0(X1)|r2_hidden(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_47, c_0_32])).
fof(c_0_52, plain, ![X20, X21]:(~r1_tarski(X20,X21)|r1_tarski(k3_tarski(X20),k3_tarski(X21))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t95_zfmisc_1])])).
cnf(c_0_53, plain, (k3_tarski(k1_zfmisc_1(X1))=X1|~r2_hidden(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_54, negated_conjecture, (r2_hidden(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_55, plain, (r1_tarski(k3_tarski(X1),k3_tarski(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_56, negated_conjecture, (k3_tarski(k1_zfmisc_1(u1_struct_0(esk1_0)))=u1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_53, c_0_46])).
cnf(c_0_57, plain, (k3_tarski(k1_zfmisc_1(X1))=X1|v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_53, c_0_51])).
cnf(c_0_58, plain, (v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|r2_hidden(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_24, c_0_23])).
cnf(c_0_59, negated_conjecture, (~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(spm,[status(thm)],[c_0_30, c_0_54])).
fof(c_0_60, plain, ![X12]:(v1_xboole_0(X12)|(~r2_hidden(k3_tarski(X12),X12)|k4_yellow_0(k2_yellow_1(X12))=k3_tarski(X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_yellow_1])])])).
cnf(c_0_61, negated_conjecture, (r1_tarski(k3_tarski(X1),u1_struct_0(esk1_0))|~r1_tarski(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_62, plain, (v1_xboole_0(X1)|r1_tarski(X2,X1)|~r2_hidden(X2,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_43, c_0_57])).
cnf(c_0_63, negated_conjecture, (r2_hidden(u1_pre_topc(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_35]), c_0_59])).
cnf(c_0_64, plain, (v1_xboole_0(X1)|k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_65, negated_conjecture, (k3_tarski(X1)=u1_struct_0(esk1_0)|~r2_hidden(u1_struct_0(esk1_0),X1)|~r1_tarski(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_48, c_0_61])).
cnf(c_0_66, negated_conjecture, (r1_tarski(u1_pre_topc(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_50])).
cnf(c_0_67, plain, (k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(csr,[status(thm)],[c_0_64, c_0_30])).
cnf(c_0_68, negated_conjecture, (k3_tarski(u1_pre_topc(esk1_0))=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_40])])).
cnf(c_0_69, negated_conjecture, (k4_yellow_0(k2_yellow_1(u1_pre_topc(esk1_0)))!=u1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_70, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_40])]), c_0_69]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 71
# Proof object clause steps            : 48
# Proof object formula steps           : 23
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 14
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 38
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 38
# Processed clauses                    : 1149
# ...of these trivial                  : 2
# ...subsumed                          : 711
# ...remaining for further processing  : 436
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 33
# Backward-rewritten                   : 57
# Generated clauses                    : 5794
# ...of the previous two non-trivial   : 5422
# Contextual simplify-reflections      : 5
# Paramodulations                      : 5786
# Factorizations                       : 6
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 307
#    Positive orientable unit clauses  : 25
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 13
#    Non-unit-clauses                  : 269
# Current number of unprocessed clauses: 4277
# ...number of literals in the above   : 17088
# Current number of archived formulas  : 0
# Current number of archived clauses   : 127
# Clause-clause subsumption calls (NU) : 7402
# Rec. Clause-clause subsumption calls : 4479
# Non-unit clause-clause subsumptions  : 521
# Unit Clause-clause subsumption calls : 74
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 19
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 82934

# -------------------------------------------------
# User time                : 0.061 s
# System time              : 0.003 s
# Total time               : 0.065 s
# Maximum resident set size: 28728 pages
