# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t30_orders_2, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~((r1_orders_2(X1,X2,X3)&r2_orders_2(X1,X3,X2)))))), file('data/MPTP/proofs/256a87461bea4baeb447167cb5fab3e6.E_input', t30_orders_2)).
fof(d10_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_orders_2(X1,X2,X3)<=>(r1_orders_2(X1,X2,X3)&~(X2=X3)))))), file('data/MPTP/proofs/256a87461bea4baeb447167cb5fab3e6.E_input', d10_orders_2)).
fof(t16_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3))))))), file('data/MPTP/proofs/256a87461bea4baeb447167cb5fab3e6.E_input', t16_yellow_0)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/256a87461bea4baeb447167cb5fab3e6.E_input', d8_lattice3)).
fof(d4_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>(v1_yellow_0(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&r1_lattice3(X1,u1_struct_0(X1),X2)))), file('data/MPTP/proofs/256a87461bea4baeb447167cb5fab3e6.E_input', d4_yellow_0)).
fof(d8_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X5)=>r1_orders_2(X1,X5,X4))))=>X4=X3))))))), file('data/MPTP/proofs/256a87461bea4baeb447167cb5fab3e6.E_input', d8_yellow_0)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/256a87461bea4baeb447167cb5fab3e6.E_input', t1_subset)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/256a87461bea4baeb447167cb5fab3e6.E_input', t2_subset)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/256a87461bea4baeb447167cb5fab3e6.E_input', d1_xboole_0)).
fof(t42_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))=>(r1_yellow_0(X1,k1_xboole_0)&r2_yellow_0(X1,u1_struct_0(X1)))), file('data/MPTP/proofs/256a87461bea4baeb447167cb5fab3e6.E_input', t42_yellow_0)).
fof(t30_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4)))))&((r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))=>(X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3)))))), file('data/MPTP/proofs/256a87461bea4baeb447167cb5fab3e6.E_input', t30_yellow_0)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/256a87461bea4baeb447167cb5fab3e6.E_input', t5_subset)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/256a87461bea4baeb447167cb5fab3e6.E_input', t4_subset_1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/256a87461bea4baeb447167cb5fab3e6.E_input', fc2_struct_0)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/256a87461bea4baeb447167cb5fab3e6.E_input', d9_lattice3)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/256a87461bea4baeb447167cb5fab3e6.E_input', dt_l1_orders_2)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/256a87461bea4baeb447167cb5fab3e6.E_input', fc1_xboole_0)).
fof(c_0_17, plain, ![X75, X76, X77]:(~v5_orders_2(X75)|~l1_orders_2(X75)|(~m1_subset_1(X76,u1_struct_0(X75))|(~m1_subset_1(X77,u1_struct_0(X75))|(~r1_orders_2(X75,X76,X77)|~r2_orders_2(X75,X77,X76))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_orders_2])])])).
fof(c_0_18, plain, ![X34, X35, X36]:(((r1_orders_2(X34,X35,X36)|~r2_orders_2(X34,X35,X36)|~m1_subset_1(X36,u1_struct_0(X34))|~m1_subset_1(X35,u1_struct_0(X34))|~l1_orders_2(X34))&(X35!=X36|~r2_orders_2(X34,X35,X36)|~m1_subset_1(X36,u1_struct_0(X34))|~m1_subset_1(X35,u1_struct_0(X34))|~l1_orders_2(X34)))&(~r1_orders_2(X34,X35,X36)|X35=X36|r2_orders_2(X34,X35,X36)|~m1_subset_1(X36,u1_struct_0(X34))|~m1_subset_1(X35,u1_struct_0(X34))|~l1_orders_2(X34))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_orders_2])])])])])).
cnf(c_0_19, plain, (~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)|~r2_orders_2(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20, plain, (X2=X3|r2_orders_2(X1,X2,X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_21, plain, ![X37, X38, X40, X41, X42]:(((m1_subset_1(esk6_2(X37,X38),u1_struct_0(X37))|~r2_yellow_0(X37,X38)|(~v5_orders_2(X37)|~l1_orders_2(X37)))&((r1_lattice3(X37,X38,esk6_2(X37,X38))|~r2_yellow_0(X37,X38)|(~v5_orders_2(X37)|~l1_orders_2(X37)))&(~m1_subset_1(X40,u1_struct_0(X37))|(~r1_lattice3(X37,X38,X40)|r1_orders_2(X37,X40,esk6_2(X37,X38)))|~r2_yellow_0(X37,X38)|(~v5_orders_2(X37)|~l1_orders_2(X37)))))&((m1_subset_1(esk7_3(X37,X41,X42),u1_struct_0(X37))|~r1_lattice3(X37,X41,X42)|~m1_subset_1(X42,u1_struct_0(X37))|r2_yellow_0(X37,X41)|(~v5_orders_2(X37)|~l1_orders_2(X37)))&((r1_lattice3(X37,X41,esk7_3(X37,X41,X42))|~r1_lattice3(X37,X41,X42)|~m1_subset_1(X42,u1_struct_0(X37))|r2_yellow_0(X37,X41)|(~v5_orders_2(X37)|~l1_orders_2(X37)))&(~r1_orders_2(X37,esk7_3(X37,X41,X42),X42)|~r1_lattice3(X37,X41,X42)|~m1_subset_1(X42,u1_struct_0(X37))|r2_yellow_0(X37,X41)|(~v5_orders_2(X37)|~l1_orders_2(X37)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_yellow_0])])])])])])).
fof(c_0_22, plain, ![X16, X17, X18, X19]:((~r1_lattice3(X16,X17,X18)|(~m1_subset_1(X19,u1_struct_0(X16))|(~r2_hidden(X19,X17)|r1_orders_2(X16,X18,X19)))|~m1_subset_1(X18,u1_struct_0(X16))|~l1_orders_2(X16))&((m1_subset_1(esk4_3(X16,X17,X18),u1_struct_0(X16))|r1_lattice3(X16,X17,X18)|~m1_subset_1(X18,u1_struct_0(X16))|~l1_orders_2(X16))&((r2_hidden(esk4_3(X16,X17,X18),X17)|r1_lattice3(X16,X17,X18)|~m1_subset_1(X18,u1_struct_0(X16))|~l1_orders_2(X16))&(~r1_orders_2(X16,X18,esk4_3(X16,X17,X18))|r1_lattice3(X16,X17,X18)|~m1_subset_1(X18,u1_struct_0(X16))|~l1_orders_2(X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
cnf(c_0_23, plain, (X1=X2|~r1_orders_2(X3,X2,X1)|~r1_orders_2(X3,X1,X2)|~m1_subset_1(X1,u1_struct_0(X3))|~m1_subset_1(X2,u1_struct_0(X3))|~l1_orders_2(X3)|~v5_orders_2(X3)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_24, plain, (r1_orders_2(X2,X1,esk6_2(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~r2_yellow_0(X2,X3)|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_25, plain, (m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_26, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27, plain, (r1_lattice3(X1,X2,esk6_2(X1,X2))|~r2_yellow_0(X1,X2)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_28, plain, ![X87, X89]:(((m1_subset_1(esk14_1(X87),u1_struct_0(X87))|~v1_yellow_0(X87)|~l1_orders_2(X87))&(r1_lattice3(X87,u1_struct_0(X87),esk14_1(X87))|~v1_yellow_0(X87)|~l1_orders_2(X87)))&(~m1_subset_1(X89,u1_struct_0(X87))|~r1_lattice3(X87,u1_struct_0(X87),X89)|v1_yellow_0(X87)|~l1_orders_2(X87))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_yellow_0])])])])])).
fof(c_0_29, plain, ![X48, X49, X51, X52, X54, X55, X58]:(((m1_subset_1(esk8_2(X48,X49),u1_struct_0(X48))|~r2_yellow_0(X48,X49)|~l1_orders_2(X48))&((r1_lattice3(X48,X49,esk8_2(X48,X49))|~r2_yellow_0(X48,X49)|~l1_orders_2(X48))&((~m1_subset_1(X51,u1_struct_0(X48))|(~r1_lattice3(X48,X49,X51)|r1_orders_2(X48,X51,esk8_2(X48,X49)))|~r2_yellow_0(X48,X49)|~l1_orders_2(X48))&((m1_subset_1(esk9_3(X48,X49,X52),u1_struct_0(X48))|~r1_lattice3(X48,X49,X52)|X52=esk8_2(X48,X49)|~m1_subset_1(X52,u1_struct_0(X48))|~r2_yellow_0(X48,X49)|~l1_orders_2(X48))&((r1_lattice3(X48,X49,esk9_3(X48,X49,X52))|~r1_lattice3(X48,X49,X52)|X52=esk8_2(X48,X49)|~m1_subset_1(X52,u1_struct_0(X48))|~r2_yellow_0(X48,X49)|~l1_orders_2(X48))&(~r1_orders_2(X48,esk9_3(X48,X49,X52),X52)|~r1_lattice3(X48,X49,X52)|X52=esk8_2(X48,X49)|~m1_subset_1(X52,u1_struct_0(X48))|~r2_yellow_0(X48,X49)|~l1_orders_2(X48)))))))&(((m1_subset_1(esk11_3(X48,X54,X55),u1_struct_0(X48))|m1_subset_1(esk10_3(X48,X54,X55),u1_struct_0(X48))|~r1_lattice3(X48,X54,X55)|~m1_subset_1(X55,u1_struct_0(X48))|r2_yellow_0(X48,X54)|~l1_orders_2(X48))&(((r1_lattice3(X48,X54,esk11_3(X48,X54,X55))|m1_subset_1(esk10_3(X48,X54,X55),u1_struct_0(X48))|~r1_lattice3(X48,X54,X55)|~m1_subset_1(X55,u1_struct_0(X48))|r2_yellow_0(X48,X54)|~l1_orders_2(X48))&(~m1_subset_1(X58,u1_struct_0(X48))|(~r1_lattice3(X48,X54,X58)|r1_orders_2(X48,X58,esk11_3(X48,X54,X55)))|m1_subset_1(esk10_3(X48,X54,X55),u1_struct_0(X48))|~r1_lattice3(X48,X54,X55)|~m1_subset_1(X55,u1_struct_0(X48))|r2_yellow_0(X48,X54)|~l1_orders_2(X48)))&(esk11_3(X48,X54,X55)!=X55|m1_subset_1(esk10_3(X48,X54,X55),u1_struct_0(X48))|~r1_lattice3(X48,X54,X55)|~m1_subset_1(X55,u1_struct_0(X48))|r2_yellow_0(X48,X54)|~l1_orders_2(X48))))&(((m1_subset_1(esk11_3(X48,X54,X55),u1_struct_0(X48))|r1_lattice3(X48,X54,esk10_3(X48,X54,X55))|~r1_lattice3(X48,X54,X55)|~m1_subset_1(X55,u1_struct_0(X48))|r2_yellow_0(X48,X54)|~l1_orders_2(X48))&(((r1_lattice3(X48,X54,esk11_3(X48,X54,X55))|r1_lattice3(X48,X54,esk10_3(X48,X54,X55))|~r1_lattice3(X48,X54,X55)|~m1_subset_1(X55,u1_struct_0(X48))|r2_yellow_0(X48,X54)|~l1_orders_2(X48))&(~m1_subset_1(X58,u1_struct_0(X48))|(~r1_lattice3(X48,X54,X58)|r1_orders_2(X48,X58,esk11_3(X48,X54,X55)))|r1_lattice3(X48,X54,esk10_3(X48,X54,X55))|~r1_lattice3(X48,X54,X55)|~m1_subset_1(X55,u1_struct_0(X48))|r2_yellow_0(X48,X54)|~l1_orders_2(X48)))&(esk11_3(X48,X54,X55)!=X55|r1_lattice3(X48,X54,esk10_3(X48,X54,X55))|~r1_lattice3(X48,X54,X55)|~m1_subset_1(X55,u1_struct_0(X48))|r2_yellow_0(X48,X54)|~l1_orders_2(X48))))&((m1_subset_1(esk11_3(X48,X54,X55),u1_struct_0(X48))|~r1_orders_2(X48,esk10_3(X48,X54,X55),X55)|~r1_lattice3(X48,X54,X55)|~m1_subset_1(X55,u1_struct_0(X48))|r2_yellow_0(X48,X54)|~l1_orders_2(X48))&(((r1_lattice3(X48,X54,esk11_3(X48,X54,X55))|~r1_orders_2(X48,esk10_3(X48,X54,X55),X55)|~r1_lattice3(X48,X54,X55)|~m1_subset_1(X55,u1_struct_0(X48))|r2_yellow_0(X48,X54)|~l1_orders_2(X48))&(~m1_subset_1(X58,u1_struct_0(X48))|(~r1_lattice3(X48,X54,X58)|r1_orders_2(X48,X58,esk11_3(X48,X54,X55)))|~r1_orders_2(X48,esk10_3(X48,X54,X55),X55)|~r1_lattice3(X48,X54,X55)|~m1_subset_1(X55,u1_struct_0(X48))|r2_yellow_0(X48,X54)|~l1_orders_2(X48)))&(esk11_3(X48,X54,X55)!=X55|~r1_orders_2(X48,esk10_3(X48,X54,X55),X55)|~r1_lattice3(X48,X54,X55)|~m1_subset_1(X55,u1_struct_0(X48))|r2_yellow_0(X48,X54)|~l1_orders_2(X48))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_yellow_0])])])])])])).
cnf(c_0_30, plain, (esk6_2(X1,X2)=X3|~r1_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk6_2(X1,X2),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
cnf(c_0_31, plain, (r1_orders_2(X1,esk6_2(X1,X2),X3)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_25])).
cnf(c_0_32, plain, (v1_yellow_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,u1_struct_0(X2),X1)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33, plain, (r1_lattice3(X1,X2,esk8_2(X1,X2))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34, plain, (m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_35, plain, ![X21, X22]:(~r2_hidden(X21,X22)|m1_subset_1(X21,X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_36, plain, (esk6_2(X1,X2)=X3|~r1_lattice3(X1,X2,X3)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_37, plain, (r1_lattice3(X1,u1_struct_0(X1),esk14_1(X1))|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, plain, (v1_yellow_0(X1)|~r2_yellow_0(X1,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])).
cnf(c_0_39, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_40, plain, ![X23, X24]:(~m1_subset_1(X23,X24)|(v1_xboole_0(X24)|r2_hidden(X23,X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_41, plain, (esk6_2(X1,u1_struct_0(X1))=esk14_1(X1)|~r2_hidden(esk14_1(X1),u1_struct_0(X1))|~r2_yellow_0(X1,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_39])).
cnf(c_0_42, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_43, plain, (esk6_2(X1,u1_struct_0(X1))=esk14_1(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(esk14_1(X1),u1_struct_0(X1))|~r2_yellow_0(X1,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_44, plain, (m1_subset_1(esk14_1(X1),u1_struct_0(X1))|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_45, plain, ![X12, X13, X14]:((~v1_xboole_0(X12)|~r2_hidden(X13,X12))&(r2_hidden(esk3_1(X14),X14)|v1_xboole_0(X14))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_46, plain, (r1_lattice3(X1,X2,esk7_3(X1,X2,X3))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_47, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_48, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))=>(r1_yellow_0(X1,k1_xboole_0)&r2_yellow_0(X1,u1_struct_0(X1))))), inference(assume_negation,[status(cth)],[t42_yellow_0])).
fof(c_0_49, plain, ![X66, X67, X68, X69, X70]:(((r2_lattice3(X66,X68,X67)|(X67!=k1_yellow_0(X66,X68)|~r1_yellow_0(X66,X68))|~m1_subset_1(X67,u1_struct_0(X66))|(~v5_orders_2(X66)|~l1_orders_2(X66)))&(~m1_subset_1(X69,u1_struct_0(X66))|(~r2_lattice3(X66,X68,X69)|r1_orders_2(X66,X67,X69))|(X67!=k1_yellow_0(X66,X68)|~r1_yellow_0(X66,X68))|~m1_subset_1(X67,u1_struct_0(X66))|(~v5_orders_2(X66)|~l1_orders_2(X66))))&(((X67=k1_yellow_0(X66,X70)|(m1_subset_1(esk12_3(X66,X67,X70),u1_struct_0(X66))|~r2_lattice3(X66,X70,X67))|~m1_subset_1(X67,u1_struct_0(X66))|(~v5_orders_2(X66)|~l1_orders_2(X66)))&(r1_yellow_0(X66,X70)|(m1_subset_1(esk12_3(X66,X67,X70),u1_struct_0(X66))|~r2_lattice3(X66,X70,X67))|~m1_subset_1(X67,u1_struct_0(X66))|(~v5_orders_2(X66)|~l1_orders_2(X66))))&(((X67=k1_yellow_0(X66,X70)|(r2_lattice3(X66,X70,esk12_3(X66,X67,X70))|~r2_lattice3(X66,X70,X67))|~m1_subset_1(X67,u1_struct_0(X66))|(~v5_orders_2(X66)|~l1_orders_2(X66)))&(r1_yellow_0(X66,X70)|(r2_lattice3(X66,X70,esk12_3(X66,X67,X70))|~r2_lattice3(X66,X70,X67))|~m1_subset_1(X67,u1_struct_0(X66))|(~v5_orders_2(X66)|~l1_orders_2(X66))))&((X67=k1_yellow_0(X66,X70)|(~r1_orders_2(X66,X67,esk12_3(X66,X67,X70))|~r2_lattice3(X66,X70,X67))|~m1_subset_1(X67,u1_struct_0(X66))|(~v5_orders_2(X66)|~l1_orders_2(X66)))&(r1_yellow_0(X66,X70)|(~r1_orders_2(X66,X67,esk12_3(X66,X67,X70))|~r2_lattice3(X66,X70,X67))|~m1_subset_1(X67,u1_struct_0(X66))|(~v5_orders_2(X66)|~l1_orders_2(X66))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_yellow_0])])])])])])).
cnf(c_0_50, plain, (esk6_2(X1,u1_struct_0(X1))=esk14_1(X1)|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_38])).
cnf(c_0_51, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_52, plain, ![X26, X27, X28]:(~r2_hidden(X26,X27)|(~m1_subset_1(X27,k1_zfmisc_1(X28))|~v1_xboole_0(X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_53, plain, ![X74]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X74)), inference(variable_rename,[status(thm)],[t4_subset_1])).
cnf(c_0_54, plain, (r2_yellow_0(X1,X2)|~r1_orders_2(X1,esk7_3(X1,X2,X3),X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_55, plain, (r1_orders_2(X1,esk7_3(X1,X2,X3),X4)|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~r2_hidden(X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_46]), c_0_47])).
fof(c_0_56, negated_conjecture, ((~v2_struct_0(esk1_0)&(v5_orders_2(esk1_0)&(v1_yellow_0(esk1_0)&l1_orders_2(esk1_0))))&(~r1_yellow_0(esk1_0,k1_xboole_0)|~r2_yellow_0(esk1_0,u1_struct_0(esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_48])])])])).
fof(c_0_57, plain, ![X61]:(v2_struct_0(X61)|~l1_struct_0(X61)|~v1_xboole_0(u1_struct_0(X61))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_58, plain, (r1_yellow_0(X1,X2)|~r1_orders_2(X1,X3,esk12_3(X1,X3,X2))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_59, plain, (r1_orders_2(X1,esk14_1(X1),X2)|~r2_hidden(X2,u1_struct_0(X1))|~r2_yellow_0(X1,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_50]), c_0_39]), c_0_51])).
cnf(c_0_60, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_61, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_53])).
fof(c_0_62, plain, ![X7, X8, X9, X10]:((~r2_lattice3(X7,X8,X9)|(~m1_subset_1(X10,u1_struct_0(X7))|(~r2_hidden(X10,X8)|r1_orders_2(X7,X10,X9)))|~m1_subset_1(X9,u1_struct_0(X7))|~l1_orders_2(X7))&((m1_subset_1(esk2_3(X7,X8,X9),u1_struct_0(X7))|r2_lattice3(X7,X8,X9)|~m1_subset_1(X9,u1_struct_0(X7))|~l1_orders_2(X7))&((r2_hidden(esk2_3(X7,X8,X9),X8)|r2_lattice3(X7,X8,X9)|~m1_subset_1(X9,u1_struct_0(X7))|~l1_orders_2(X7))&(~r1_orders_2(X7,esk2_3(X7,X8,X9),X9)|r2_lattice3(X7,X8,X9)|~m1_subset_1(X9,u1_struct_0(X7))|~l1_orders_2(X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
cnf(c_0_63, plain, (r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_64, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_65, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_57])).
fof(c_0_66, plain, ![X25]:(~l1_orders_2(X25)|l1_struct_0(X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_67, plain, (r1_yellow_0(X1,X2)|~r2_hidden(esk12_3(X1,esk14_1(X1),X2),u1_struct_0(X1))|~r2_lattice3(X1,X2,esk14_1(X1))|~m1_subset_1(esk14_1(X1),u1_struct_0(X1))|~r2_yellow_0(X1,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_68, plain, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(esk14_1(X1),u1_struct_0(X1))|~r2_yellow_0(X1,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_25, c_0_50])).
cnf(c_0_69, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_70, plain, (r2_hidden(esk2_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_71, plain, (r2_yellow_0(X1,u1_struct_0(X1))|~r2_hidden(esk14_1(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_37]), c_0_39])).
cnf(c_0_72, negated_conjecture, (~l1_struct_0(esk1_0)|~v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_73, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_74, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_75, plain, (v1_xboole_0(u1_struct_0(X1))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,esk14_1(X1))|~m1_subset_1(esk12_3(X1,esk14_1(X1),X2),u1_struct_0(X1))|~r2_yellow_0(X1,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_42]), c_0_68])).
cnf(c_0_76, plain, (r1_yellow_0(X1,X2)|m1_subset_1(esk12_3(X1,X3,X2),u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_77, plain, (r2_lattice3(X1,k1_xboole_0,X2)|~v1_xboole_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_78, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_79, plain, (v1_xboole_0(u1_struct_0(X1))|r2_yellow_0(X1,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_42]), c_0_44])).
cnf(c_0_80, negated_conjecture, (v1_yellow_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_81, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_82, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_74])])).
cnf(c_0_83, plain, (v1_xboole_0(u1_struct_0(X1))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,esk14_1(X1))|~r2_yellow_0(X1,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_68])).
cnf(c_0_84, plain, (r2_lattice3(X1,k1_xboole_0,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_85, negated_conjecture, (~r1_yellow_0(esk1_0,k1_xboole_0)|~r2_yellow_0(esk1_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_86, negated_conjecture, (r2_yellow_0(esk1_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_74]), c_0_81])]), c_0_82])).
cnf(c_0_87, plain, (v1_xboole_0(u1_struct_0(X1))|r1_yellow_0(X1,k1_xboole_0)|~r2_yellow_0(X1,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_68])).
cnf(c_0_88, negated_conjecture, (~r1_yellow_0(esk1_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_86])])).
cnf(c_0_89, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_86]), c_0_74]), c_0_81])]), c_0_82]), c_0_88]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 90
# Proof object clause steps            : 56
# Proof object formula steps           : 34
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 30
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 27
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 27
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 86
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 86
# Processed clauses                    : 34306
# ...of these trivial                  : 0
# ...subsumed                          : 28565
# ...remaining for further processing  : 5741
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1770
# Backward-rewritten                   : 1238
# Generated clauses                    : 170066
# ...of the previous two non-trivial   : 163824
# Contextual simplify-reflections      : 1243
# Paramodulations                      : 170057
# Factorizations                       : 2
# Equation resolutions                 : 7
# Propositional unsat checks           : 1
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.160
#    Propositional solver time         : 0.040
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2644
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 2629
# Current number of unprocessed clauses: 125700
# ...number of literals in the above   : 1093145
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3090
# Clause-clause subsumption calls (NU) : 3237747
# Rec. Clause-clause subsumption calls : 194686
# Non-unit clause-clause subsumptions  : 29435
# Unit Clause-clause subsumption calls : 7372
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6128572

# -------------------------------------------------
# User time                : 4.874 s
# System time              : 0.060 s
# Total time               : 4.934 s
# Maximum resident set size: 139204 pages
