# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t5_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1)&v2_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1)))), file('data/MPTP/proofs/2580aff522d1436b8ddd3fcd66c43017.E_input', t5_waybel_0)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/2580aff522d1436b8ddd3fcd66c43017.E_input', t5_subset)).
fof(t38_orders_2, axiom, ![X1]:((v3_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(~((?[X4]:((v6_orders_2(X4,X1)&m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1))))&(r2_hidden(X2,X4)&r2_hidden(X3,X4)))&(~(r1_orders_2(X1,X2,X3))&~(r1_orders_2(X1,X3,X2)))))&~(((r1_orders_2(X1,X2,X3)|r1_orders_2(X1,X3,X2))&![X4]:((v6_orders_2(X4,X1)&m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1))))=>~((r2_hidden(X2,X4)&r2_hidden(X3,X4)))))))))), file('data/MPTP/proofs/2580aff522d1436b8ddd3fcd66c43017.E_input', t38_orders_2)).
fof(reflexivity_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>r3_orders_2(X1,X2,X2)), file('data/MPTP/proofs/2580aff522d1436b8ddd3fcd66c43017.E_input', reflexivity_r3_orders_2)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/2580aff522d1436b8ddd3fcd66c43017.E_input', d1_tarski)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/2580aff522d1436b8ddd3fcd66c43017.E_input', t69_enumset1)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/2580aff522d1436b8ddd3fcd66c43017.E_input', redefinition_k6_domain_1)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/2580aff522d1436b8ddd3fcd66c43017.E_input', t24_orders_2)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/2580aff522d1436b8ddd3fcd66c43017.E_input', redefinition_r3_orders_2)).
fof(d2_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v2_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X3,X2)&(r2_hidden(X4,X2)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>~((r2_hidden(X5,X2)&(r1_orders_2(X1,X5,X3)&r1_orders_2(X1,X5,X4))))))))))))), file('data/MPTP/proofs/2580aff522d1436b8ddd3fcd66c43017.E_input', d2_waybel_0)).
fof(d1_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v1_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X3,X2)&(r2_hidden(X4,X2)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>~((r2_hidden(X5,X2)&(r1_orders_2(X1,X3,X5)&r1_orders_2(X1,X4,X5))))))))))))), file('data/MPTP/proofs/2580aff522d1436b8ddd3fcd66c43017.E_input', d1_waybel_0)).
fof(dt_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))), file('data/MPTP/proofs/2580aff522d1436b8ddd3fcd66c43017.E_input', dt_k6_domain_1)).
fof(c_0_12, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1)&v2_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1))))), inference(assume_negation,[status(cth)],[t5_waybel_0])).
fof(c_0_13, plain, ![X8, X9, X10]:(~r2_hidden(X8,X9)|(~m1_subset_1(X9,k1_zfmisc_1(X10))|~v1_xboole_0(X10))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_14, plain, ![X39, X40, X41, X42]:((~v6_orders_2(X42,X39)|~m1_subset_1(X42,k1_zfmisc_1(u1_struct_0(X39)))|(~r2_hidden(X40,X42)|~r2_hidden(X41,X42))|(r1_orders_2(X39,X40,X41)|r1_orders_2(X39,X41,X40))|~m1_subset_1(X41,u1_struct_0(X39))|~m1_subset_1(X40,u1_struct_0(X39))|(~v3_orders_2(X39)|~l1_orders_2(X39)))&((((v6_orders_2(esk7_3(X39,X40,X41),X39)|~r1_orders_2(X39,X40,X41)|~m1_subset_1(X41,u1_struct_0(X39))|~m1_subset_1(X40,u1_struct_0(X39))|(~v3_orders_2(X39)|~l1_orders_2(X39)))&(m1_subset_1(esk7_3(X39,X40,X41),k1_zfmisc_1(u1_struct_0(X39)))|~r1_orders_2(X39,X40,X41)|~m1_subset_1(X41,u1_struct_0(X39))|~m1_subset_1(X40,u1_struct_0(X39))|(~v3_orders_2(X39)|~l1_orders_2(X39))))&((r2_hidden(X40,esk7_3(X39,X40,X41))|~r1_orders_2(X39,X40,X41)|~m1_subset_1(X41,u1_struct_0(X39))|~m1_subset_1(X40,u1_struct_0(X39))|(~v3_orders_2(X39)|~l1_orders_2(X39)))&(r2_hidden(X41,esk7_3(X39,X40,X41))|~r1_orders_2(X39,X40,X41)|~m1_subset_1(X41,u1_struct_0(X39))|~m1_subset_1(X40,u1_struct_0(X39))|(~v3_orders_2(X39)|~l1_orders_2(X39)))))&(((v6_orders_2(esk7_3(X39,X40,X41),X39)|~r1_orders_2(X39,X41,X40)|~m1_subset_1(X41,u1_struct_0(X39))|~m1_subset_1(X40,u1_struct_0(X39))|(~v3_orders_2(X39)|~l1_orders_2(X39)))&(m1_subset_1(esk7_3(X39,X40,X41),k1_zfmisc_1(u1_struct_0(X39)))|~r1_orders_2(X39,X41,X40)|~m1_subset_1(X41,u1_struct_0(X39))|~m1_subset_1(X40,u1_struct_0(X39))|(~v3_orders_2(X39)|~l1_orders_2(X39))))&((r2_hidden(X40,esk7_3(X39,X40,X41))|~r1_orders_2(X39,X41,X40)|~m1_subset_1(X41,u1_struct_0(X39))|~m1_subset_1(X40,u1_struct_0(X39))|(~v3_orders_2(X39)|~l1_orders_2(X39)))&(r2_hidden(X41,esk7_3(X39,X40,X41))|~r1_orders_2(X39,X41,X40)|~m1_subset_1(X41,u1_struct_0(X39))|~m1_subset_1(X40,u1_struct_0(X39))|(~v3_orders_2(X39)|~l1_orders_2(X39))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t38_orders_2])])])])])])).
fof(c_0_15, plain, ![X27, X28, X29]:(v2_struct_0(X27)|(~v3_orders_2(X27)|~l1_orders_2(X27))|(~m1_subset_1(X28,u1_struct_0(X27))|~m1_subset_1(X29,u1_struct_0(X27)))|r3_orders_2(X27,X28,X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r3_orders_2])])])).
fof(c_0_16, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&l1_orders_2(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(~v1_waybel_0(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)|~v2_waybel_0(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])).
fof(c_0_17, plain, ![X30, X31, X32, X33, X34, X35]:(((~r2_hidden(X32,X31)|X32=X30|X31!=k1_tarski(X30))&(X33!=X30|r2_hidden(X33,X31)|X31!=k1_tarski(X30)))&((~r2_hidden(esk6_2(X34,X35),X35)|esk6_2(X34,X35)!=X34|X35=k1_tarski(X34))&(r2_hidden(esk6_2(X34,X35),X35)|esk6_2(X34,X35)=X34|X35=k1_tarski(X34)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_18, plain, ![X21]:k2_tarski(X21,X21)=k1_tarski(X21), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_19, plain, ![X22, X23]:(v1_xboole_0(X22)|~m1_subset_1(X23,X22)|k6_domain_1(X22,X23)=k1_tarski(X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_20, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (m1_subset_1(esk7_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_22, plain, ![X37, X38]:(v2_struct_0(X37)|(~v3_orders_2(X37)|~l1_orders_2(X37))|(~m1_subset_1(X38,u1_struct_0(X37))|r1_orders_2(X37,X38,X38))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_23, plain, ![X24, X25, X26]:((~r3_orders_2(X24,X25,X26)|r1_orders_2(X24,X25,X26)|(v2_struct_0(X24)|(~v3_orders_2(X24)|~l1_orders_2(X24))|(~m1_subset_1(X25,u1_struct_0(X24))|~m1_subset_1(X26,u1_struct_0(X24)))))&(~r1_orders_2(X24,X25,X26)|r3_orders_2(X24,X25,X26)|(v2_struct_0(X24)|(~v3_orders_2(X24)|~l1_orders_2(X24))|(~m1_subset_1(X25,u1_struct_0(X24))|~m1_subset_1(X26,u1_struct_0(X24)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_24, plain, (v2_struct_0(X1)|r3_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_27, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_29, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_30, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_31, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_32, plain, (~r1_orders_2(X1,X2,X3)|~v1_xboole_0(u1_struct_0(X1))|~r2_hidden(X4,esk7_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_33, plain, (r2_hidden(X1,esk7_3(X2,X3,X1))|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~v3_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_34, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_35, plain, ![X13, X14, X15, X16, X20]:(((m1_subset_1(esk3_4(X13,X14,X15,X16),u1_struct_0(X13))|~r2_hidden(X16,X14)|~r2_hidden(X15,X14)|~m1_subset_1(X16,u1_struct_0(X13))|~m1_subset_1(X15,u1_struct_0(X13))|~v2_waybel_0(X14,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))&((r2_hidden(esk3_4(X13,X14,X15,X16),X14)|~r2_hidden(X16,X14)|~r2_hidden(X15,X14)|~m1_subset_1(X16,u1_struct_0(X13))|~m1_subset_1(X15,u1_struct_0(X13))|~v2_waybel_0(X14,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))&((r1_orders_2(X13,esk3_4(X13,X14,X15,X16),X15)|~r2_hidden(X16,X14)|~r2_hidden(X15,X14)|~m1_subset_1(X16,u1_struct_0(X13))|~m1_subset_1(X15,u1_struct_0(X13))|~v2_waybel_0(X14,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))&(r1_orders_2(X13,esk3_4(X13,X14,X15,X16),X16)|~r2_hidden(X16,X14)|~r2_hidden(X15,X14)|~m1_subset_1(X16,u1_struct_0(X13))|~m1_subset_1(X15,u1_struct_0(X13))|~v2_waybel_0(X14,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13)))))&((m1_subset_1(esk4_2(X13,X14),u1_struct_0(X13))|v2_waybel_0(X14,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))&((m1_subset_1(esk5_2(X13,X14),u1_struct_0(X13))|v2_waybel_0(X14,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))&((r2_hidden(esk4_2(X13,X14),X14)|v2_waybel_0(X14,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))&((r2_hidden(esk5_2(X13,X14),X14)|v2_waybel_0(X14,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))&(~m1_subset_1(X20,u1_struct_0(X13))|(~r2_hidden(X20,X14)|(~r1_orders_2(X13,X20,esk4_2(X13,X14))|~r1_orders_2(X13,X20,esk5_2(X13,X14))))|v2_waybel_0(X14,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_waybel_0])])])])])).
cnf(c_0_36, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_37, negated_conjecture, (r3_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26]), c_0_27])]), c_0_28])).
cnf(c_0_38, plain, (X1=X3|X2!=k2_tarski(X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_39, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_40, plain, (k2_tarski(X2,X2)=k6_domain_1(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_31, c_0_30])).
cnf(c_0_41, plain, (~r1_orders_2(X1,X2,X3)|~v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_42, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_25]), c_0_26]), c_0_27])]), c_0_28])).
cnf(c_0_43, plain, (v2_waybel_0(X3,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r1_orders_2(X2,X1,esk4_2(X2,X3))|~r1_orders_2(X2,X1,esk5_2(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_44, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_26]), c_0_27])]), c_0_28])).
cnf(c_0_45, plain, (X1=X2|~r2_hidden(X1,k2_tarski(X2,X2))), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_46, plain, (r2_hidden(esk5_2(X1,X2),X2)|v2_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_47, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X2)), inference(rw,[status(thm)],[c_0_39, c_0_30])).
fof(c_0_48, plain, ![X44, X45, X46, X47, X51]:(((m1_subset_1(esk8_4(X44,X45,X46,X47),u1_struct_0(X44))|~r2_hidden(X47,X45)|~r2_hidden(X46,X45)|~m1_subset_1(X47,u1_struct_0(X44))|~m1_subset_1(X46,u1_struct_0(X44))|~v1_waybel_0(X45,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X44)))|~l1_orders_2(X44))&((r2_hidden(esk8_4(X44,X45,X46,X47),X45)|~r2_hidden(X47,X45)|~r2_hidden(X46,X45)|~m1_subset_1(X47,u1_struct_0(X44))|~m1_subset_1(X46,u1_struct_0(X44))|~v1_waybel_0(X45,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X44)))|~l1_orders_2(X44))&((r1_orders_2(X44,X46,esk8_4(X44,X45,X46,X47))|~r2_hidden(X47,X45)|~r2_hidden(X46,X45)|~m1_subset_1(X47,u1_struct_0(X44))|~m1_subset_1(X46,u1_struct_0(X44))|~v1_waybel_0(X45,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X44)))|~l1_orders_2(X44))&(r1_orders_2(X44,X47,esk8_4(X44,X45,X46,X47))|~r2_hidden(X47,X45)|~r2_hidden(X46,X45)|~m1_subset_1(X47,u1_struct_0(X44))|~m1_subset_1(X46,u1_struct_0(X44))|~v1_waybel_0(X45,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X44)))|~l1_orders_2(X44)))))&((m1_subset_1(esk9_2(X44,X45),u1_struct_0(X44))|v1_waybel_0(X45,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X44)))|~l1_orders_2(X44))&((m1_subset_1(esk10_2(X44,X45),u1_struct_0(X44))|v1_waybel_0(X45,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X44)))|~l1_orders_2(X44))&((r2_hidden(esk9_2(X44,X45),X45)|v1_waybel_0(X45,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X44)))|~l1_orders_2(X44))&((r2_hidden(esk10_2(X44,X45),X45)|v1_waybel_0(X45,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X44)))|~l1_orders_2(X44))&(~m1_subset_1(X51,u1_struct_0(X44))|(~r2_hidden(X51,X45)|(~r1_orders_2(X44,esk9_2(X44,X45),X51)|~r1_orders_2(X44,esk10_2(X44,X45),X51)))|v1_waybel_0(X45,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X44)))|~l1_orders_2(X44))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_waybel_0])])])])])).
cnf(c_0_49, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_40, c_0_25])).
cnf(c_0_50, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_25]), c_0_26]), c_0_27])])).
cnf(c_0_51, negated_conjecture, (v2_waybel_0(X1,esk1_0)|~r1_orders_2(esk1_0,esk5_2(esk1_0,X1),esk4_2(esk1_0,X1))|~r2_hidden(esk5_2(esk1_0,X1),X1)|~m1_subset_1(esk5_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_26])])).
cnf(c_0_52, plain, (esk5_2(X1,k2_tarski(X2,X2))=X2|v2_waybel_0(k2_tarski(X2,X2),X1)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_53, plain, (r2_hidden(X1,k2_tarski(X1,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_47])])).
cnf(c_0_54, plain, (r2_hidden(esk4_2(X1,X2),X2)|v2_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_55, plain, ![X11, X12]:(v1_xboole_0(X11)|~m1_subset_1(X12,X11)|m1_subset_1(k6_domain_1(X11,X12),k1_zfmisc_1(X11))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_domain_1])])])).
cnf(c_0_56, plain, (v1_waybel_0(X3,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r1_orders_2(X2,esk9_2(X2,X3),X1)|~r1_orders_2(X2,esk10_2(X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_57, plain, (r2_hidden(esk10_2(X1,X2),X2)|v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_58, negated_conjecture, (~v1_waybel_0(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)|~v2_waybel_0(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_59, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)), inference(sr,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_60, negated_conjecture, (v2_waybel_0(k2_tarski(X1,X1),esk1_0)|~r1_orders_2(esk1_0,X1,esk4_2(esk1_0,k2_tarski(X1,X1)))|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53]), c_0_26])])).
cnf(c_0_61, plain, (esk4_2(X1,k2_tarski(X2,X2))=X2|v2_waybel_0(k2_tarski(X2,X2),X1)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_45, c_0_54])).
cnf(c_0_62, plain, (v1_xboole_0(X1)|m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_63, negated_conjecture, (v1_waybel_0(X1,esk1_0)|~r1_orders_2(esk1_0,esk9_2(esk1_0,X1),esk10_2(esk1_0,X1))|~r2_hidden(esk10_2(esk1_0,X1),X1)|~m1_subset_1(esk10_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_44]), c_0_26])])).
cnf(c_0_64, plain, (esk10_2(X1,k2_tarski(X2,X2))=X2|v1_waybel_0(k2_tarski(X2,X2),X1)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_45, c_0_57])).
cnf(c_0_65, plain, (r2_hidden(esk9_2(X1,X2),X2)|v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_66, negated_conjecture, (~v2_waybel_0(k2_tarski(esk2_0,esk2_0),esk1_0)|~v1_waybel_0(k2_tarski(esk2_0,esk2_0),esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_59]), c_0_59])).
cnf(c_0_67, negated_conjecture, (v2_waybel_0(k2_tarski(X1,X1),esk1_0)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_26])]), c_0_44])).
cnf(c_0_68, negated_conjecture, (m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_59]), c_0_25])]), c_0_50])).
cnf(c_0_69, negated_conjecture, (v1_waybel_0(k2_tarski(X1,X1),esk1_0)|~r1_orders_2(esk1_0,esk9_2(esk1_0,k2_tarski(X1,X1)),X1)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_53]), c_0_26])])).
cnf(c_0_70, plain, (esk9_2(X1,k2_tarski(X2,X2))=X2|v1_waybel_0(k2_tarski(X2,X2),X1)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_45, c_0_65])).
cnf(c_0_71, negated_conjecture, (~v1_waybel_0(k2_tarski(esk2_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68]), c_0_25])])).
cnf(c_0_72, negated_conjecture, (v1_waybel_0(k2_tarski(X1,X1),esk1_0)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_26])]), c_0_44])).
cnf(c_0_73, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_68]), c_0_25])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 74
# Proof object clause steps            : 49
# Proof object formula steps           : 25
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 50
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 44
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 43
# Processed clauses                    : 143
# ...of these trivial                  : 0
# ...subsumed                          : 15
# ...remaining for further processing  : 128
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 2
# Generated clauses                    : 297
# ...of the previous two non-trivial   : 285
# Contextual simplify-reflections      : 2
# Paramodulations                      : 293
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 77
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 67
# Current number of unprocessed clauses: 222
# ...number of literals in the above   : 1542
# Current number of archived formulas  : 0
# Current number of archived clauses   : 50
# Clause-clause subsumption calls (NU) : 2322
# Rec. Clause-clause subsumption calls : 211
# Non-unit clause-clause subsumptions  : 18
# Unit Clause-clause subsumption calls : 27
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 12319

# -------------------------------------------------
# User time                : 0.018 s
# System time              : 0.000 s
# Total time               : 0.018 s
# Maximum resident set size: 28672 pages
