# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', t12_setfam_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', t70_enumset1)).
fof(t25_pre_topc, conjecture, ![X1]:(l1_struct_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((k2_struct_0(X1)=k4_subset_1(u1_struct_0(X1),X2,X3)&r1_xboole_0(X2,X3))=>X3=k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))))), file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', t25_pre_topc)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', t100_xboole_1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', l51_zfmisc_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', d7_xboole_0)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', symmetry_r1_xboole_0)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', commutativity_k2_tarski)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', t3_boole)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', t2_boole)).
fof(redefinition_k4_subset_1, axiom, ![X1, X2, X3]:((m1_subset_1(X2,k1_zfmisc_1(X1))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>k4_subset_1(X1,X2,X3)=k2_xboole_0(X2,X3)), file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', redefinition_k4_subset_1)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', d3_struct_0)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', t98_xboole_1)).
fof(t46_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0, file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', t46_xboole_1)).
fof(t21_xboole_1, axiom, ![X1, X2]:k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', t21_xboole_1)).
fof(redefinition_k7_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k7_subset_1(X1,X2,X3)=k4_xboole_0(X2,X3)), file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', redefinition_k7_subset_1)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', d4_subset_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', t36_xboole_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', idempotence_k3_xboole_0)).
fof(t42_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k4_xboole_0(X1,X3),k4_xboole_0(X2,X3)), file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', t42_xboole_1)).
fof(t45_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>X2=k2_xboole_0(X1,k4_xboole_0(X2,X1))), file('data/MPTP/proofs/25851078d97146c09fa71961c8568ac8.E_input', t45_xboole_1)).
fof(c_0_22, plain, ![X25, X26]:k1_setfam_1(k2_tarski(X25,X26))=k3_xboole_0(X25,X26), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_23, plain, ![X44, X45]:k1_enumset1(X44,X44,X45)=k2_tarski(X44,X45), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_24, negated_conjecture, ~(![X1]:(l1_struct_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((k2_struct_0(X1)=k4_subset_1(u1_struct_0(X1),X2,X3)&r1_xboole_0(X2,X3))=>X3=k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2)))))), inference(assume_negation,[status(cth)],[t25_pre_topc])).
fof(c_0_25, plain, ![X21, X22]:k4_xboole_0(X21,X22)=k5_xboole_0(X21,k3_xboole_0(X21,X22)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_26, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_28, plain, ![X14, X15]:k3_tarski(k2_tarski(X14,X15))=k2_xboole_0(X14,X15), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_29, plain, ![X12, X13]:((~r1_xboole_0(X12,X13)|k3_xboole_0(X12,X13)=k1_xboole_0)&(k3_xboole_0(X12,X13)!=k1_xboole_0|r1_xboole_0(X12,X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
fof(c_0_30, plain, ![X23, X24]:(~r1_xboole_0(X23,X24)|r1_xboole_0(X24,X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
fof(c_0_31, negated_conjecture, (l1_struct_0(esk1_0)&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&((k2_struct_0(esk1_0)=k4_subset_1(u1_struct_0(esk1_0),esk2_0,esk3_0)&r1_xboole_0(esk2_0,esk3_0))&esk3_0!=k7_subset_1(u1_struct_0(esk1_0),k2_struct_0(esk1_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).
fof(c_0_32, plain, ![X17, X18]:k2_tarski(X17,X18)=k2_tarski(X18,X17), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
fof(c_0_33, plain, ![X7]:k4_xboole_0(X7,k1_xboole_0)=X7, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_34, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_26, c_0_27])).
fof(c_0_36, plain, ![X19]:k3_xboole_0(X19,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_37, plain, ![X34, X35, X36]:(~m1_subset_1(X35,k1_zfmisc_1(X34))|~m1_subset_1(X36,k1_zfmisc_1(X34))|k4_subset_1(X34,X35,X36)=k2_xboole_0(X35,X36)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_subset_1])])).
cnf(c_0_38, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_39, plain, ![X33]:(~l1_struct_0(X33)|k2_struct_0(X33)=u1_struct_0(X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
fof(c_0_40, plain, ![X40, X41]:k2_xboole_0(X40,X41)=k5_xboole_0(X40,k4_xboole_0(X41,X40)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
cnf(c_0_41, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_42, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_43, negated_conjecture, (r1_xboole_0(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_44, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_45, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_46, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_47, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_48, plain, ![X37, X38]:k4_xboole_0(X37,k2_xboole_0(X37,X38))=k1_xboole_0, inference(variable_rename,[status(thm)],[t46_xboole_1])).
fof(c_0_49, plain, ![X42, X43]:k3_xboole_0(X42,k2_xboole_0(X42,X43))=X42, inference(variable_rename,[status(thm)],[t21_xboole_1])).
cnf(c_0_50, plain, (k4_subset_1(X2,X1,X3)=k2_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_51, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_38, c_0_27])).
cnf(c_0_52, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_53, negated_conjecture, (l1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_54, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_55, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_41, c_0_35])).
cnf(c_0_56, negated_conjecture, (r1_xboole_0(esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_57, plain, (k1_enumset1(X1,X1,X2)=k1_enumset1(X2,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_27]), c_0_27])).
cnf(c_0_58, plain, (k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,k1_xboole_0)))=X1), inference(rw,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_59, plain, (k1_setfam_1(k1_enumset1(X1,X1,k1_xboole_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_47, c_0_35])).
fof(c_0_60, plain, ![X27, X28, X29]:(~m1_subset_1(X28,k1_zfmisc_1(X27))|k7_subset_1(X27,X28,X29)=k4_xboole_0(X28,X29)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_subset_1])])).
fof(c_0_61, plain, ![X16]:m1_subset_1(k2_subset_1(X16),k1_zfmisc_1(X16)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_62, plain, ![X39]:k2_subset_1(X39)=X39, inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_63, plain, ![X8, X9]:r1_tarski(k4_xboole_0(X8,X9),X8), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_64, plain, ![X20]:k3_xboole_0(X20,X20)=X20, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_65, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_66, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_67, plain, (k4_subset_1(X2,X1,X3)=k3_tarski(k1_enumset1(X1,X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_68, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_69, negated_conjecture, (k2_struct_0(esk1_0)=k4_subset_1(u1_struct_0(esk1_0),esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_70, negated_conjecture, (k2_struct_0(esk1_0)=u1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_71, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k5_xboole_0(X1,k5_xboole_0(X2,k1_setfam_1(k1_enumset1(X2,X2,X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_51]), c_0_46])).
cnf(c_0_72, negated_conjecture, (k1_setfam_1(k1_enumset1(esk2_0,esk2_0,esk3_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57])).
cnf(c_0_73, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_58, c_0_59])).
fof(c_0_74, plain, ![X30, X31, X32]:k4_xboole_0(k2_xboole_0(X30,X31),X32)=k2_xboole_0(k4_xboole_0(X30,X32),k4_xboole_0(X31,X32)), inference(variable_rename,[status(thm)],[t42_xboole_1])).
cnf(c_0_75, plain, (k7_subset_1(X2,X1,X3)=k4_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_76, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_77, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_62])).
fof(c_0_78, plain, ![X10, X11]:(~r1_tarski(X10,X11)|X11=k2_xboole_0(X10,k4_xboole_0(X11,X10))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t45_xboole_1])])).
cnf(c_0_79, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_80, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_81, plain, (k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,k3_tarski(k1_enumset1(X1,X1,X2)))))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_51]), c_0_46])).
cnf(c_0_82, plain, (k1_setfam_1(k1_enumset1(X1,X1,k3_tarski(k1_enumset1(X1,X1,X2))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_51]), c_0_35])).
cnf(c_0_83, negated_conjecture, (esk3_0!=k7_subset_1(u1_struct_0(esk1_0),k2_struct_0(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_84, negated_conjecture, (k4_subset_1(u1_struct_0(esk1_0),X1,esk3_0)=k3_tarski(k1_enumset1(X1,X1,esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_85, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_86, negated_conjecture, (k4_subset_1(u1_struct_0(esk1_0),esk2_0,esk3_0)=u1_struct_0(esk1_0)), inference(rw,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_87, negated_conjecture, (k3_tarski(k1_enumset1(esk2_0,esk2_0,esk3_0))=k5_xboole_0(esk3_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73]), c_0_57])).
cnf(c_0_88, plain, (k4_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k4_xboole_0(X1,X3),k4_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_89, plain, (k7_subset_1(X2,X1,X3)=k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X3)))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_75, c_0_46])).
cnf(c_0_90, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_91, plain, (X2=k2_xboole_0(X1,k4_xboole_0(X2,X1))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_92, plain, (r1_tarski(k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X2))),X1)), inference(rw,[status(thm)],[c_0_79, c_0_46])).
cnf(c_0_93, plain, (k1_setfam_1(k1_enumset1(X1,X1,X1))=X1), inference(rw,[status(thm)],[c_0_80, c_0_35])).
cnf(c_0_94, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_95, negated_conjecture, (k7_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),esk2_0)!=esk3_0), inference(rw,[status(thm)],[c_0_83, c_0_70])).
cnf(c_0_96, negated_conjecture, (u1_struct_0(esk1_0)=k5_xboole_0(esk3_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86]), c_0_87])).
cnf(c_0_97, plain, (k5_xboole_0(X1,k5_xboole_0(X2,k1_setfam_1(k1_enumset1(X1,X1,X2))))=k3_tarski(k1_enumset1(X1,X1,X2))), inference(spm,[status(thm)],[c_0_71, c_0_57])).
cnf(c_0_98, plain, (k5_xboole_0(k3_tarski(k1_enumset1(X1,X1,X2)),k1_setfam_1(k1_enumset1(k3_tarski(k1_enumset1(X1,X1,X2)),k3_tarski(k1_enumset1(X1,X1,X2)),X3)))=k3_tarski(k1_enumset1(k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X3))),k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X3))),k5_xboole_0(X2,k1_setfam_1(k1_enumset1(X2,X2,X3)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_88, c_0_51]), c_0_51]), c_0_46]), c_0_46]), c_0_46]), c_0_46])).
cnf(c_0_99, plain, (k7_subset_1(X1,X1,X2)=k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X2)))), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_100, plain, (X2=k3_tarski(k1_enumset1(X1,X1,k5_xboole_0(X2,k1_setfam_1(k1_enumset1(X2,X2,X1)))))|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_51]), c_0_46])).
cnf(c_0_101, plain, (r1_tarski(k1_xboole_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_94])).
cnf(c_0_102, negated_conjecture, (k7_subset_1(k5_xboole_0(esk3_0,esk2_0),k5_xboole_0(esk3_0,esk2_0),esk2_0)!=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_95, c_0_96]), c_0_96])).
cnf(c_0_103, negated_conjecture, (k5_xboole_0(esk3_0,esk2_0)=k5_xboole_0(esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_72]), c_0_73]), c_0_87])).
cnf(c_0_104, negated_conjecture, (k1_setfam_1(k1_enumset1(esk2_0,esk2_0,k5_xboole_0(esk3_0,esk2_0)))=esk2_0), inference(spm,[status(thm)],[c_0_82, c_0_87])).
cnf(c_0_105, plain, (k3_tarski(k1_enumset1(k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X2))),k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X2))),k5_xboole_0(X3,k1_setfam_1(k1_enumset1(X3,X3,X2)))))=k7_subset_1(k3_tarski(k1_enumset1(X1,X1,X3)),k3_tarski(k1_enumset1(X1,X1,X3)),X2)), inference(rw,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_106, plain, (k3_tarski(k1_enumset1(k1_xboole_0,k1_xboole_0,X1))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_59]), c_0_73]), c_0_101])])).
cnf(c_0_107, negated_conjecture, (k7_subset_1(k5_xboole_0(esk2_0,esk3_0),k5_xboole_0(esk2_0,esk3_0),esk2_0)!=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_102, c_0_103]), c_0_103])).
cnf(c_0_108, plain, (k7_subset_1(X1,X1,X2)=k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X2,X2,X1)))), inference(spm,[status(thm)],[c_0_99, c_0_57])).
cnf(c_0_109, negated_conjecture, (k1_setfam_1(k1_enumset1(esk2_0,esk2_0,k5_xboole_0(esk2_0,esk3_0)))=esk2_0), inference(rw,[status(thm)],[c_0_104, c_0_103])).
cnf(c_0_110, plain, (k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X2)))=k7_subset_1(k3_tarski(k1_enumset1(X2,X2,X1)),k3_tarski(k1_enumset1(X2,X2,X1)),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_93]), c_0_94]), c_0_94]), c_0_106])).
cnf(c_0_111, negated_conjecture, (k3_tarski(k1_enumset1(esk2_0,esk2_0,esk3_0))=k5_xboole_0(esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_87, c_0_103])).
cnf(c_0_112, negated_conjecture, (k5_xboole_0(k5_xboole_0(esk2_0,esk3_0),esk2_0)!=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_107, c_0_108]), c_0_109])).
cnf(c_0_113, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_57]), c_0_72]), c_0_73]), c_0_108]), c_0_109]), c_0_112]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 114
# Proof object clause steps            : 69
# Proof object formula steps           : 45
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 22
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 60
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 23
# Processed clauses                    : 688
# ...of these trivial                  : 56
# ...subsumed                          : 346
# ...remaining for further processing  : 286
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 61
# Generated clauses                    : 3035
# ...of the previous two non-trivial   : 2126
# Contextual simplify-reflections      : 1
# Paramodulations                      : 3027
# Factorizations                       : 0
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 199
#    Positive orientable unit clauses  : 86
#    Positive unorientable unit clauses: 7
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 104
# Current number of unprocessed clauses: 1380
# ...number of literals in the above   : 2164
# Current number of archived formulas  : 0
# Current number of archived clauses   : 92
# Clause-clause subsumption calls (NU) : 3548
# Rec. Clause-clause subsumption calls : 3287
# Non-unit clause-clause subsumptions  : 337
# Unit Clause-clause subsumption calls : 277
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 432
# BW rewrite match successes           : 57
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 51312

# -------------------------------------------------
# User time                : 0.030 s
# System time              : 0.003 s
# Total time               : 0.033 s
# Maximum resident set size: 28616 pages
