# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.007 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t32_mcart_1, conjecture, ![X1, X2, X3, X4]:k4_mcart_1(X1,X2,X3,X4)=k3_mcart_1(k4_tarski(X1,X2),X3,X4), file('data/MPTP/proofs/25932b8100144a35aa926bd98186c870.E_input', t32_mcart_1)).
fof(d1_mcart_1, axiom, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>![X2]:(X2=k1_mcart_1(X1)<=>![X3, X4]:(X1=k4_tarski(X3,X4)=>X2=X3))), file('data/MPTP/proofs/25932b8100144a35aa926bd98186c870.E_input', d1_mcart_1)).
fof(d4_mcart_1, axiom, ![X1, X2, X3, X4]:k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4), file('data/MPTP/proofs/25932b8100144a35aa926bd98186c870.E_input', d4_mcart_1)).
fof(t31_mcart_1, axiom, ![X1, X2, X3, X4]:k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4), file('data/MPTP/proofs/25932b8100144a35aa926bd98186c870.E_input', t31_mcart_1)).
fof(c_0_4, negated_conjecture, ~(![X1, X2, X3, X4]:k4_mcart_1(X1,X2,X3,X4)=k3_mcart_1(k4_tarski(X1,X2),X3,X4)), inference(assume_negation,[status(cth)],[t32_mcart_1])).
fof(c_0_5, plain, ![X17, X18, X19, X20, X21, X22, X23]:((X20!=k1_mcart_1(X17)|(X17!=k4_tarski(X21,X22)|X20=X21)|X17!=k4_tarski(X18,X19))&((X17=k4_tarski(esk5_2(X17,X23),esk6_2(X17,X23))|X23=k1_mcart_1(X17)|X17!=k4_tarski(X18,X19))&(X23!=esk5_2(X17,X23)|X23=k1_mcart_1(X17)|X17!=k4_tarski(X18,X19)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_mcart_1])])])])])])).
fof(c_0_6, negated_conjecture, k4_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)!=k3_mcart_1(k4_tarski(esk1_0,esk2_0),esk3_0,esk4_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
fof(c_0_7, plain, ![X9, X10, X11, X12]:k4_mcart_1(X9,X10,X11,X12)=k4_tarski(k3_mcart_1(X9,X10,X11),X12), inference(variable_rename,[status(thm)],[d4_mcart_1])).
fof(c_0_8, plain, ![X13, X14, X15, X16]:k4_mcart_1(X13,X14,X15,X16)=k4_tarski(k4_tarski(k4_tarski(X13,X14),X15),X16), inference(variable_rename,[status(thm)],[t31_mcart_1])).
cnf(c_0_9, plain, (X1=X3|X1!=k1_mcart_1(X2)|X2!=k4_tarski(X3,X4)|X2!=k4_tarski(X5,X6)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10, negated_conjecture, (k4_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)!=k3_mcart_1(k4_tarski(esk1_0,esk2_0),esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11, plain, (k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (k1_mcart_1(k4_tarski(X1,X2))=X3|k4_tarski(X1,X2)!=k4_tarski(X3,X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_9])])).
cnf(c_0_14, negated_conjecture, (k4_tarski(k3_mcart_1(esk1_0,esk2_0,esk3_0),esk4_0)!=k3_mcart_1(k4_tarski(esk1_0,esk2_0),esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15, plain, (k4_tarski(k3_mcart_1(X1,X2,X3),X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_12, c_0_11])).
cnf(c_0_16, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(er,[status(thm)],[c_0_13])).
cnf(c_0_17, negated_conjecture, (k3_mcart_1(k4_tarski(esk1_0,esk2_0),esk3_0,esk4_0)!=k4_tarski(k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0),esk4_0)), inference(rw,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_18, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_15]), c_0_16])).
cnf(c_0_19, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_18])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 20
# Proof object clause steps            : 11
# Proof object formula steps           : 9
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 4
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 2
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 6
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 5
# Processed clauses                    : 18
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 18
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 5
# Generated clauses                    : 19
# ...of the previous two non-trivial   : 21
# Contextual simplify-reflections      : 0
# Paramodulations                      : 14
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 5
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 3
# Current number of unprocessed clauses: 13
# ...number of literals in the above   : 30
# Current number of archived formulas  : 0
# Current number of archived clauses   : 11
# Clause-clause subsumption calls (NU) : 4
# Rec. Clause-clause subsumption calls : 4
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 566

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.000 s
# Total time               : 0.008 s
# Maximum resident set size: 28440 pages
