# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t40_funct_1, conjecture, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k1_relat_1(k5_relat_1(X3,k6_relat_1(X2))))<=>(r2_hidden(X1,k1_relat_1(X3))&r2_hidden(k1_funct_1(X3,X1),X2)))), file('data/MPTP/proofs/25a0174abe6745079efddb3522b3a38c.E_input', t40_funct_1)).
fof(t21_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k1_relat_1(k5_relat_1(X3,X2)))<=>(r2_hidden(X1,k1_relat_1(X3))&r2_hidden(k1_funct_1(X3,X1),k1_relat_1(X2)))))), file('data/MPTP/proofs/25a0174abe6745079efddb3522b3a38c.E_input', t21_funct_1)).
fof(fc3_funct_1, axiom, ![X1]:(v1_relat_1(k6_relat_1(X1))&v1_funct_1(k6_relat_1(X1))), file('data/MPTP/proofs/25a0174abe6745079efddb3522b3a38c.E_input', fc3_funct_1)).
fof(d10_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(X2=k6_relat_1(X1)<=>![X3, X4]:(r2_hidden(k4_tarski(X3,X4),X2)<=>(r2_hidden(X3,X1)&X3=X4)))), file('data/MPTP/proofs/25a0174abe6745079efddb3522b3a38c.E_input', d10_relat_1)).
fof(d4_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2, X3]:((r2_hidden(X2,k1_relat_1(X1))=>(X3=k1_funct_1(X1,X2)<=>r2_hidden(k4_tarski(X2,X3),X1)))&(~(r2_hidden(X2,k1_relat_1(X1)))=>(X3=k1_funct_1(X1,X2)<=>X3=k1_xboole_0)))), file('data/MPTP/proofs/25a0174abe6745079efddb3522b3a38c.E_input', d4_funct_1)).
fof(t8_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>(r2_hidden(X1,k1_relat_1(X3))&X2=k1_funct_1(X3,X1)))), file('data/MPTP/proofs/25a0174abe6745079efddb3522b3a38c.E_input', t8_funct_1)).
fof(c_0_6, negated_conjecture, ~(![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k1_relat_1(k5_relat_1(X3,k6_relat_1(X2))))<=>(r2_hidden(X1,k1_relat_1(X3))&r2_hidden(k1_funct_1(X3,X1),X2))))), inference(assume_negation,[status(cth)],[t40_funct_1])).
fof(c_0_7, plain, ![X23, X24, X25]:(((r2_hidden(X23,k1_relat_1(X25))|~r2_hidden(X23,k1_relat_1(k5_relat_1(X25,X24)))|(~v1_relat_1(X25)|~v1_funct_1(X25))|(~v1_relat_1(X24)|~v1_funct_1(X24)))&(r2_hidden(k1_funct_1(X25,X23),k1_relat_1(X24))|~r2_hidden(X23,k1_relat_1(k5_relat_1(X25,X24)))|(~v1_relat_1(X25)|~v1_funct_1(X25))|(~v1_relat_1(X24)|~v1_funct_1(X24))))&(~r2_hidden(X23,k1_relat_1(X25))|~r2_hidden(k1_funct_1(X25,X23),k1_relat_1(X24))|r2_hidden(X23,k1_relat_1(k5_relat_1(X25,X24)))|(~v1_relat_1(X25)|~v1_funct_1(X25))|(~v1_relat_1(X24)|~v1_funct_1(X24)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t21_funct_1])])])])).
fof(c_0_8, negated_conjecture, ((v1_relat_1(esk3_0)&v1_funct_1(esk3_0))&((~r2_hidden(esk1_0,k1_relat_1(k5_relat_1(esk3_0,k6_relat_1(esk2_0))))|(~r2_hidden(esk1_0,k1_relat_1(esk3_0))|~r2_hidden(k1_funct_1(esk3_0,esk1_0),esk2_0)))&((r2_hidden(esk1_0,k1_relat_1(esk3_0))|r2_hidden(esk1_0,k1_relat_1(k5_relat_1(esk3_0,k6_relat_1(esk2_0)))))&(r2_hidden(k1_funct_1(esk3_0,esk1_0),esk2_0)|r2_hidden(esk1_0,k1_relat_1(k5_relat_1(esk3_0,k6_relat_1(esk2_0)))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
fof(c_0_9, plain, ![X11]:(v1_relat_1(k6_relat_1(X11))&v1_funct_1(k6_relat_1(X11))), inference(variable_rename,[status(thm)],[fc3_funct_1])).
fof(c_0_10, plain, ![X15, X16, X17, X18, X19, X20]:((((r2_hidden(X17,X15)|~r2_hidden(k4_tarski(X17,X18),X16)|X16!=k6_relat_1(X15)|~v1_relat_1(X16))&(X17=X18|~r2_hidden(k4_tarski(X17,X18),X16)|X16!=k6_relat_1(X15)|~v1_relat_1(X16)))&(~r2_hidden(X19,X15)|X19!=X20|r2_hidden(k4_tarski(X19,X20),X16)|X16!=k6_relat_1(X15)|~v1_relat_1(X16)))&((~r2_hidden(k4_tarski(esk4_2(X15,X16),esk5_2(X15,X16)),X16)|(~r2_hidden(esk4_2(X15,X16),X15)|esk4_2(X15,X16)!=esk5_2(X15,X16))|X16=k6_relat_1(X15)|~v1_relat_1(X16))&((r2_hidden(esk4_2(X15,X16),X15)|r2_hidden(k4_tarski(esk4_2(X15,X16),esk5_2(X15,X16)),X16)|X16=k6_relat_1(X15)|~v1_relat_1(X16))&(esk4_2(X15,X16)=esk5_2(X15,X16)|r2_hidden(k4_tarski(esk4_2(X15,X16),esk5_2(X15,X16)),X16)|X16=k6_relat_1(X15)|~v1_relat_1(X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_relat_1])])])])])])).
fof(c_0_11, plain, ![X8, X9, X10]:(((X10!=k1_funct_1(X8,X9)|r2_hidden(k4_tarski(X9,X10),X8)|~r2_hidden(X9,k1_relat_1(X8))|(~v1_relat_1(X8)|~v1_funct_1(X8)))&(~r2_hidden(k4_tarski(X9,X10),X8)|X10=k1_funct_1(X8,X9)|~r2_hidden(X9,k1_relat_1(X8))|(~v1_relat_1(X8)|~v1_funct_1(X8))))&((X10!=k1_funct_1(X8,X9)|X10=k1_xboole_0|r2_hidden(X9,k1_relat_1(X8))|(~v1_relat_1(X8)|~v1_funct_1(X8)))&(X10!=k1_xboole_0|X10=k1_funct_1(X8,X9)|r2_hidden(X9,k1_relat_1(X8))|(~v1_relat_1(X8)|~v1_funct_1(X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_funct_1])])])])])).
cnf(c_0_12, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(X1,k1_relat_1(k5_relat_1(X2,X3)))|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(esk3_0))|r2_hidden(esk1_0,k1_relat_1(k5_relat_1(esk3_0,k6_relat_1(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, plain, (v1_funct_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16, plain, (v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_18, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),X4)|X4!=k6_relat_1(X2)|~v1_relat_1(X4)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19, plain, (r2_hidden(k4_tarski(X3,X1),X2)|X1!=k1_funct_1(X2,X3)|~r2_hidden(X3,k1_relat_1(X2))|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20, negated_conjecture, (~r2_hidden(esk1_0,k1_relat_1(k5_relat_1(esk3_0,k6_relat_1(esk2_0))))|~r2_hidden(esk1_0,k1_relat_1(esk3_0))|~r2_hidden(k1_funct_1(esk3_0,esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_21, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_15]), c_0_16]), c_0_17])])).
cnf(c_0_22, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k6_relat_1(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_18]), c_0_16])])).
cnf(c_0_23, plain, (r2_hidden(k4_tarski(X1,k1_funct_1(X2,X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_19])).
fof(c_0_24, plain, ![X12, X13, X14]:(((r2_hidden(X12,k1_relat_1(X14))|~r2_hidden(k4_tarski(X12,X13),X14)|(~v1_relat_1(X14)|~v1_funct_1(X14)))&(X13=k1_funct_1(X14,X12)|~r2_hidden(k4_tarski(X12,X13),X14)|(~v1_relat_1(X14)|~v1_funct_1(X14))))&(~r2_hidden(X12,k1_relat_1(X14))|X13!=k1_funct_1(X14,X12)|r2_hidden(k4_tarski(X12,X13),X14)|(~v1_relat_1(X14)|~v1_funct_1(X14)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_funct_1])])])).
cnf(c_0_25, plain, (r2_hidden(k4_tarski(X1,X3),X4)|~r2_hidden(X1,X2)|X1!=X3|X4!=k6_relat_1(X2)|~v1_relat_1(X4)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_26, plain, (r2_hidden(k1_funct_1(X1,X2),k1_relat_1(X3))|~r2_hidden(X2,k1_relat_1(k5_relat_1(X1,X3)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_27, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk1_0),esk2_0)|r2_hidden(esk1_0,k1_relat_1(k5_relat_1(esk3_0,k6_relat_1(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_28, negated_conjecture, (~r2_hidden(esk1_0,k1_relat_1(k5_relat_1(esk3_0,k6_relat_1(esk2_0))))|~r2_hidden(k1_funct_1(esk3_0,esk1_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_21])])).
cnf(c_0_29, plain, (r2_hidden(X1,k1_relat_1(k5_relat_1(X2,X3)))|~r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k1_funct_1(X2,X1),k1_relat_1(X3))|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_30, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_relat_1(k6_relat_1(X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_14]), c_0_16])])).
cnf(c_0_31, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, plain, (r2_hidden(k4_tarski(X1,X1),k6_relat_1(X2))|~r2_hidden(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_25])]), c_0_16])])).
cnf(c_0_33, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk1_0),k1_relat_1(k6_relat_1(esk2_0)))|r2_hidden(k1_funct_1(esk3_0,esk1_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_14]), c_0_15]), c_0_16]), c_0_17])])).
cnf(c_0_34, negated_conjecture, (~r2_hidden(k1_funct_1(esk3_0,esk1_0),k1_relat_1(k6_relat_1(esk2_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_21]), c_0_14]), c_0_15]), c_0_16]), c_0_17])]), c_0_30])).
cnf(c_0_35, plain, (r2_hidden(X1,k1_relat_1(k6_relat_1(X2)))|~r2_hidden(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_14]), c_0_16])])).
cnf(c_0_36, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk1_0),esk2_0)), inference(sr,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_37, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 38
# Proof object clause steps            : 25
# Proof object formula steps           : 13
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 36
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 24
# Processed clauses                    : 60
# ...of these trivial                  : 2
# ...subsumed                          : 4
# ...remaining for further processing  : 54
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 4
# Generated clauses                    : 43
# ...of the previous two non-trivial   : 31
# Contextual simplify-reflections      : 1
# Paramodulations                      : 35
# Factorizations                       : 0
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 22
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 14
# Current number of unprocessed clauses: 10
# ...number of literals in the above   : 60
# Current number of archived formulas  : 0
# Current number of archived clauses   : 25
# Clause-clause subsumption calls (NU) : 330
# Rec. Clause-clause subsumption calls : 43
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 9
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2698

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.003 s
# Total time               : 0.010 s
# Maximum resident set size: 28580 pages
