# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t175_relat_1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>k10_relat_1(X3,k2_xboole_0(X1,X2))=k2_xboole_0(k10_relat_1(X3,X1),k10_relat_1(X3,X2))), file('data/MPTP/proofs/2614479885bf421f8f04a7f73440dc29.E_input', t175_relat_1)).
fof(d3_xboole_0, axiom, ![X1, X2, X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2)))), file('data/MPTP/proofs/2614479885bf421f8f04a7f73440dc29.E_input', d3_xboole_0)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/2614479885bf421f8f04a7f73440dc29.E_input', t98_xboole_1)).
fof(d14_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(X3=k10_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(k4_tarski(X4,X5),X1)&r2_hidden(X5,X2))))), file('data/MPTP/proofs/2614479885bf421f8f04a7f73440dc29.E_input', d14_relat_1)).
fof(t166_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k10_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k2_relat_1(X3))&(r2_hidden(k4_tarski(X1,X4),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/2614479885bf421f8f04a7f73440dc29.E_input', t166_relat_1)).
fof(c_0_5, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>k10_relat_1(X3,k2_xboole_0(X1,X2))=k2_xboole_0(k10_relat_1(X3,X1),k10_relat_1(X3,X2)))), inference(assume_negation,[status(cth)],[t175_relat_1])).
fof(c_0_6, plain, ![X28, X29, X30, X31, X32, X33, X34, X35]:(((~r2_hidden(X31,X30)|(r2_hidden(X31,X28)|r2_hidden(X31,X29))|X30!=k2_xboole_0(X28,X29))&((~r2_hidden(X32,X28)|r2_hidden(X32,X30)|X30!=k2_xboole_0(X28,X29))&(~r2_hidden(X32,X29)|r2_hidden(X32,X30)|X30!=k2_xboole_0(X28,X29))))&(((~r2_hidden(esk8_3(X33,X34,X35),X33)|~r2_hidden(esk8_3(X33,X34,X35),X35)|X35=k2_xboole_0(X33,X34))&(~r2_hidden(esk8_3(X33,X34,X35),X34)|~r2_hidden(esk8_3(X33,X34,X35),X35)|X35=k2_xboole_0(X33,X34)))&(r2_hidden(esk8_3(X33,X34,X35),X35)|(r2_hidden(esk8_3(X33,X34,X35),X33)|r2_hidden(esk8_3(X33,X34,X35),X34))|X35=k2_xboole_0(X33,X34)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])).
fof(c_0_7, plain, ![X21, X22]:k2_xboole_0(X21,X22)=k5_xboole_0(X21,k4_xboole_0(X22,X21)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_8, negated_conjecture, (v1_relat_1(esk3_0)&k10_relat_1(esk3_0,k2_xboole_0(esk1_0,esk2_0))!=k2_xboole_0(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
cnf(c_0_9, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X2!=k2_xboole_0(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, negated_conjecture, (k10_relat_1(esk3_0,k2_xboole_0(esk1_0,esk2_0))!=k2_xboole_0(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (r2_hidden(esk8_3(X1,X2,X3),X3)|r2_hidden(esk8_3(X1,X2,X3),X1)|r2_hidden(esk8_3(X1,X2,X3),X2)|X3=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_13, plain, ![X9, X10, X11, X12, X14, X15, X16, X17, X19]:((((r2_hidden(k4_tarski(X12,esk4_4(X9,X10,X11,X12)),X9)|~r2_hidden(X12,X11)|X11!=k10_relat_1(X9,X10)|~v1_relat_1(X9))&(r2_hidden(esk4_4(X9,X10,X11,X12),X10)|~r2_hidden(X12,X11)|X11!=k10_relat_1(X9,X10)|~v1_relat_1(X9)))&(~r2_hidden(k4_tarski(X14,X15),X9)|~r2_hidden(X15,X10)|r2_hidden(X14,X11)|X11!=k10_relat_1(X9,X10)|~v1_relat_1(X9)))&((~r2_hidden(esk5_3(X9,X16,X17),X17)|(~r2_hidden(k4_tarski(esk5_3(X9,X16,X17),X19),X9)|~r2_hidden(X19,X16))|X17=k10_relat_1(X9,X16)|~v1_relat_1(X9))&((r2_hidden(k4_tarski(esk5_3(X9,X16,X17),esk6_3(X9,X16,X17)),X9)|r2_hidden(esk5_3(X9,X16,X17),X17)|X17=k10_relat_1(X9,X16)|~v1_relat_1(X9))&(r2_hidden(esk6_3(X9,X16,X17),X16)|r2_hidden(esk5_3(X9,X16,X17),X17)|X17=k10_relat_1(X9,X16)|~v1_relat_1(X9))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d14_relat_1])])])])])])).
cnf(c_0_14, plain, (r2_hidden(X1,X4)|r2_hidden(X1,X3)|X2!=k5_xboole_0(X3,k4_xboole_0(X4,X3))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_9, c_0_10])).
fof(c_0_15, plain, ![X23, X24, X25, X27]:(((r2_hidden(esk7_3(X23,X24,X25),k2_relat_1(X25))|~r2_hidden(X23,k10_relat_1(X25,X24))|~v1_relat_1(X25))&((r2_hidden(k4_tarski(X23,esk7_3(X23,X24,X25)),X25)|~r2_hidden(X23,k10_relat_1(X25,X24))|~v1_relat_1(X25))&(r2_hidden(esk7_3(X23,X24,X25),X24)|~r2_hidden(X23,k10_relat_1(X25,X24))|~v1_relat_1(X25))))&(~r2_hidden(X27,k2_relat_1(X25))|(~r2_hidden(k4_tarski(X23,X27),X25)|~r2_hidden(X27,X24))|r2_hidden(X23,k10_relat_1(X25,X24))|~v1_relat_1(X25))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t166_relat_1])])])])])).
cnf(c_0_16, negated_conjecture, (k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))!=k5_xboole_0(k10_relat_1(esk3_0,esk1_0),k4_xboole_0(k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,esk1_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_10]), c_0_10])).
cnf(c_0_17, plain, (X3=k5_xboole_0(X1,k4_xboole_0(X2,X1))|r2_hidden(esk8_3(X1,X2,X3),X3)|r2_hidden(esk8_3(X1,X2,X3),X2)|r2_hidden(esk8_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_12, c_0_10])).
cnf(c_0_18, plain, (r2_hidden(X1,X5)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X2,X4)|X5!=k10_relat_1(X3,X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k5_xboole_0(X2,k4_xboole_0(X3,X2)))), inference(er,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (r2_hidden(esk7_3(X1,X2,X3),X2)|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, negated_conjecture, (r2_hidden(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),X1),k10_relat_1(esk3_0,esk1_0))|r2_hidden(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),X1),k10_relat_1(esk3_0,esk2_0))|r2_hidden(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),X1),X1)|X1!=k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_22, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_23, plain, (r2_hidden(X1,k10_relat_1(X2,X3))|~r2_hidden(k4_tarski(X1,X4),X2)|~r2_hidden(X4,X3)|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (r2_hidden(k4_tarski(X1,esk7_3(X1,X2,X3)),X3)|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25, plain, (r2_hidden(esk7_3(X1,k5_xboole_0(X2,k4_xboole_0(X3,X2)),X4),X3)|r2_hidden(esk7_3(X1,k5_xboole_0(X2,k4_xboole_0(X3,X2)),X4),X2)|~r2_hidden(X1,k10_relat_1(X4,k5_xboole_0(X2,k4_xboole_0(X3,X2))))|~v1_relat_1(X4)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_26, negated_conjecture, (r2_hidden(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0))))|r2_hidden(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),k10_relat_1(esk3_0,esk2_0))|r2_hidden(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),k10_relat_1(esk3_0,esk1_0))), inference(er,[status(thm)],[c_0_21])).
cnf(c_0_27, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_28, plain, (r2_hidden(X1,X3)|X3!=k5_xboole_0(X4,k4_xboole_0(X2,X4))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_22, c_0_10])).
cnf(c_0_29, plain, (r2_hidden(X1,k10_relat_1(X2,X3))|~r2_hidden(esk7_3(X1,X4,X2),X3)|~r2_hidden(X1,k10_relat_1(X2,X4))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_30, negated_conjecture, (r2_hidden(esk7_3(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)),esk3_0),esk1_0)|r2_hidden(esk7_3(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)),esk3_0),esk2_0)|r2_hidden(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),k10_relat_1(esk3_0,esk1_0))|r2_hidden(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),k10_relat_1(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])])).
cnf(c_0_31, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_32, plain, (r2_hidden(X1,k5_xboole_0(X2,k4_xboole_0(X3,X2)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_28])).
cnf(c_0_33, negated_conjecture, (r2_hidden(esk7_3(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)),esk3_0),esk1_0)|r2_hidden(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),k10_relat_1(esk3_0,esk1_0))|r2_hidden(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),k10_relat_1(esk3_0,esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_27])]), c_0_26])).
cnf(c_0_34, plain, (r2_hidden(k4_tarski(X1,esk4_4(X2,X3,X4,X1)),X2)|~r2_hidden(X1,X4)|X4!=k10_relat_1(X2,X3)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_35, plain, (r2_hidden(X1,X3)|X3!=k5_xboole_0(X2,k4_xboole_0(X4,X2))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_31, c_0_10])).
cnf(c_0_36, plain, (r2_hidden(X1,k10_relat_1(X2,k5_xboole_0(X3,k4_xboole_0(X4,X3))))|~r2_hidden(esk7_3(X1,X5,X2),X4)|~r2_hidden(X1,k10_relat_1(X2,X5))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_29, c_0_32])).
cnf(c_0_37, negated_conjecture, (r2_hidden(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),k10_relat_1(esk3_0,esk2_0))|r2_hidden(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),k10_relat_1(esk3_0,esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_33]), c_0_27])]), c_0_26])).
cnf(c_0_38, plain, (r2_hidden(X1,k10_relat_1(X2,X3))|X4!=k10_relat_1(X2,X5)|~r2_hidden(esk4_4(X2,X5,X4,X1),X3)|~r2_hidden(X1,X4)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_23, c_0_34])).
cnf(c_0_39, plain, (r2_hidden(X1,k5_xboole_0(X2,k4_xboole_0(X3,X2)))|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_35])).
cnf(c_0_40, plain, (X3=k2_xboole_0(X1,X2)|~r2_hidden(esk8_3(X1,X2,X3),X2)|~r2_hidden(esk8_3(X1,X2,X3),X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_41, negated_conjecture, (r2_hidden(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),k10_relat_1(esk3_0,k5_xboole_0(X1,k4_xboole_0(X2,X1))))|r2_hidden(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),k10_relat_1(esk3_0,esk1_0))|~r2_hidden(esk7_3(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),esk2_0,esk3_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_27])])).
cnf(c_0_42, plain, (r2_hidden(X1,k10_relat_1(X2,k5_xboole_0(X3,k4_xboole_0(X4,X3))))|X5!=k10_relat_1(X2,X6)|~r2_hidden(esk4_4(X2,X6,X5,X1),X3)|~r2_hidden(X1,X5)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_43, plain, (X3=k5_xboole_0(X1,k4_xboole_0(X2,X1))|~r2_hidden(esk8_3(X1,X2,X3),X3)|~r2_hidden(esk8_3(X1,X2,X3),X2)), inference(rw,[status(thm)],[c_0_40, c_0_10])).
cnf(c_0_44, negated_conjecture, (r2_hidden(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),k10_relat_1(esk3_0,k5_xboole_0(X1,k4_xboole_0(esk2_0,X1))))|r2_hidden(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),k10_relat_1(esk3_0,esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_20]), c_0_27])]), c_0_37])).
cnf(c_0_45, plain, (r2_hidden(X1,k10_relat_1(X2,k5_xboole_0(X3,k4_xboole_0(X4,X3))))|~r2_hidden(esk4_4(X2,X5,k10_relat_1(X2,X5),X1),X3)|~r2_hidden(X1,k10_relat_1(X2,X5))|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),k10_relat_1(esk3_0,esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_16]), c_0_37])).
cnf(c_0_47, plain, (X3=k2_xboole_0(X1,X2)|~r2_hidden(esk8_3(X1,X2,X3),X1)|~r2_hidden(esk8_3(X1,X2,X3),X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_48, negated_conjecture, (r2_hidden(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),k10_relat_1(esk3_0,k5_xboole_0(X1,k4_xboole_0(X2,X1))))|~r2_hidden(esk4_4(esk3_0,esk1_0,k10_relat_1(esk3_0,esk1_0),esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0))))),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_27])])).
cnf(c_0_49, plain, (r2_hidden(esk4_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k10_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_50, plain, (X3=k5_xboole_0(X1,k4_xboole_0(X2,X1))|~r2_hidden(esk8_3(X1,X2,X3),X3)|~r2_hidden(esk8_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_47, c_0_10])).
cnf(c_0_51, negated_conjecture, (r2_hidden(esk8_3(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))),k10_relat_1(esk3_0,k5_xboole_0(esk1_0,k4_xboole_0(X1,esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_46]), c_0_27])])).
cnf(c_0_52, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_46])]), c_0_16]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 53
# Proof object clause steps            : 42
# Proof object formula steps           : 11
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 31
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 18
# Processed clauses                    : 1390
# ...of these trivial                  : 39
# ...subsumed                          : 757
# ...remaining for further processing  : 594
# Other redundant clauses eliminated   : 156
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 159
# Backward-rewritten                   : 47
# Generated clauses                    : 30357
# ...of the previous two non-trivial   : 29348
# Contextual simplify-reflections      : 33
# Paramodulations                      : 29941
# Factorizations                       : 198
# Equation resolutions                 : 218
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 370
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 359
# Current number of unprocessed clauses: 27817
# ...number of literals in the above   : 198809
# Current number of archived formulas  : 0
# Current number of archived clauses   : 225
# Clause-clause subsumption calls (NU) : 176620
# Rec. Clause-clause subsumption calls : 15471
# Non-unit clause-clause subsumptions  : 947
# Unit Clause-clause subsumption calls : 2133
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 71
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1140075

# -------------------------------------------------
# User time                : 0.635 s
# System time              : 0.023 s
# Total time               : 0.659 s
# Maximum resident set size: 30892 pages
