# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t14_ordinal1, conjecture, ![X1]:~(X1=k1_ordinal1(X1)), file('data/MPTP/proofs/26e87ffb12d643feb4df41723e6a9b2b.E_input', t14_ordinal1)).
fof(d1_ordinal1, axiom, ![X1]:k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1)), file('data/MPTP/proofs/26e87ffb12d643feb4df41723e6a9b2b.E_input', d1_ordinal1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/26e87ffb12d643feb4df41723e6a9b2b.E_input', t69_enumset1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/26e87ffb12d643feb4df41723e6a9b2b.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/26e87ffb12d643feb4df41723e6a9b2b.E_input', t70_enumset1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/26e87ffb12d643feb4df41723e6a9b2b.E_input', t7_ordinal1)).
fof(l49_zfmisc_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>r1_tarski(X1,k3_tarski(X2))), file('data/MPTP/proofs/26e87ffb12d643feb4df41723e6a9b2b.E_input', l49_zfmisc_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/26e87ffb12d643feb4df41723e6a9b2b.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/26e87ffb12d643feb4df41723e6a9b2b.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/26e87ffb12d643feb4df41723e6a9b2b.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/26e87ffb12d643feb4df41723e6a9b2b.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/26e87ffb12d643feb4df41723e6a9b2b.E_input', t75_enumset1)).
fof(d6_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:(X9=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)<=>![X10]:(r2_hidden(X10,X9)<=>~((~(X10=X1)&(~(X10=X2)&(~(X10=X3)&(~(X10=X4)&(~(X10=X5)&(~(X10=X6)&(~(X10=X7)&~(X10=X8))))))))))), file('data/MPTP/proofs/26e87ffb12d643feb4df41723e6a9b2b.E_input', d6_enumset1)).
fof(c_0_13, negated_conjecture, ~(![X1]:~(X1=k1_ordinal1(X1))), inference(assume_negation,[status(cth)],[t14_ordinal1])).
fof(c_0_14, plain, ![X12]:k1_ordinal1(X12)=k2_xboole_0(X12,k1_tarski(X12)), inference(variable_rename,[status(thm)],[d1_ordinal1])).
fof(c_0_15, plain, ![X13]:k2_tarski(X13,X13)=k1_tarski(X13), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_16, plain, ![X14, X15]:k3_tarski(k2_tarski(X14,X15))=k2_xboole_0(X14,X15), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_17, plain, ![X16, X17]:k1_enumset1(X16,X16,X17)=k2_tarski(X16,X17), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_18, plain, ![X18, X19]:(~r2_hidden(X18,X19)|~r1_tarski(X19,X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_19, plain, ![X20, X21]:(~r2_hidden(X20,X21)|r1_tarski(X20,k3_tarski(X21))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l49_zfmisc_1])])).
fof(c_0_20, negated_conjecture, esk1_0=k1_ordinal1(esk1_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_13])])])])).
cnf(c_0_21, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_25, plain, ![X22, X23, X24]:k2_enumset1(X22,X22,X23,X24)=k1_enumset1(X22,X23,X24), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_26, plain, ![X25, X26, X27, X28]:k3_enumset1(X25,X25,X26,X27,X28)=k2_enumset1(X25,X26,X27,X28), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_27, plain, ![X29, X30, X31, X32, X33]:k4_enumset1(X29,X29,X30,X31,X32,X33)=k3_enumset1(X29,X30,X31,X32,X33), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_28, plain, ![X34, X35, X36, X37, X38, X39]:k5_enumset1(X34,X34,X35,X36,X37,X38,X39)=k4_enumset1(X34,X35,X36,X37,X38,X39), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_29, plain, ![X40, X41, X42, X43, X44, X45, X46]:k6_enumset1(X40,X40,X41,X42,X43,X44,X45,X46)=k5_enumset1(X40,X41,X42,X43,X44,X45,X46), inference(variable_rename,[status(thm)],[t75_enumset1])).
cnf(c_0_30, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_31, plain, (r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_32, negated_conjecture, (esk1_0=k1_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_33, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k2_tarski(X1,X1))), inference(rw,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_34, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_35, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_36, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_37, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_38, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_39, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_40, plain, ![X47, X48, X49, X50, X51, X52, X53, X54, X55, X56, X57, X58, X59, X60, X61, X62, X63, X64, X65, X66]:(((~r2_hidden(X56,X55)|(X56=X47|(X56=X48|(X56=X49|(X56=X50|(X56=X51|(X56=X52|(X56=X53|X56=X54)))))))|X55!=k6_enumset1(X47,X48,X49,X50,X51,X52,X53,X54))&((X57!=X47|r2_hidden(X57,X55)|X55!=k6_enumset1(X47,X48,X49,X50,X51,X52,X53,X54))&((X57!=X48|r2_hidden(X57,X55)|X55!=k6_enumset1(X47,X48,X49,X50,X51,X52,X53,X54))&((X57!=X49|r2_hidden(X57,X55)|X55!=k6_enumset1(X47,X48,X49,X50,X51,X52,X53,X54))&((X57!=X50|r2_hidden(X57,X55)|X55!=k6_enumset1(X47,X48,X49,X50,X51,X52,X53,X54))&((X57!=X51|r2_hidden(X57,X55)|X55!=k6_enumset1(X47,X48,X49,X50,X51,X52,X53,X54))&((X57!=X52|r2_hidden(X57,X55)|X55!=k6_enumset1(X47,X48,X49,X50,X51,X52,X53,X54))&((X57!=X53|r2_hidden(X57,X55)|X55!=k6_enumset1(X47,X48,X49,X50,X51,X52,X53,X54))&(X57!=X54|r2_hidden(X57,X55)|X55!=k6_enumset1(X47,X48,X49,X50,X51,X52,X53,X54))))))))))&(((esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66)!=X58|~r2_hidden(esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66),X66)|X66=k6_enumset1(X58,X59,X60,X61,X62,X63,X64,X65))&((esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66)!=X59|~r2_hidden(esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66),X66)|X66=k6_enumset1(X58,X59,X60,X61,X62,X63,X64,X65))&((esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66)!=X60|~r2_hidden(esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66),X66)|X66=k6_enumset1(X58,X59,X60,X61,X62,X63,X64,X65))&((esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66)!=X61|~r2_hidden(esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66),X66)|X66=k6_enumset1(X58,X59,X60,X61,X62,X63,X64,X65))&((esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66)!=X62|~r2_hidden(esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66),X66)|X66=k6_enumset1(X58,X59,X60,X61,X62,X63,X64,X65))&((esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66)!=X63|~r2_hidden(esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66),X66)|X66=k6_enumset1(X58,X59,X60,X61,X62,X63,X64,X65))&((esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66)!=X64|~r2_hidden(esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66),X66)|X66=k6_enumset1(X58,X59,X60,X61,X62,X63,X64,X65))&(esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66)!=X65|~r2_hidden(esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66),X66)|X66=k6_enumset1(X58,X59,X60,X61,X62,X63,X64,X65)))))))))&(r2_hidden(esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66),X66)|(esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66)=X58|(esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66)=X59|(esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66)=X60|(esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66)=X61|(esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66)=X62|(esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66)=X63|(esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66)=X64|esk2_9(X58,X59,X60,X61,X62,X63,X64,X65,X66)=X65)))))))|X66=k6_enumset1(X58,X59,X60,X61,X62,X63,X64,X65)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_enumset1])])])])])])])).
cnf(c_0_41, plain, (~r2_hidden(k3_tarski(X1),X2)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_42, negated_conjecture, (esk1_0=k3_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_33]), c_0_24]), c_0_34]), c_0_35]), c_0_35]), c_0_36]), c_0_36]), c_0_37]), c_0_37]), c_0_38]), c_0_38]), c_0_39]), c_0_39])).
cnf(c_0_43, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k6_enumset1(X4,X5,X6,X7,X8,X9,X10,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_44, negated_conjecture, (~r2_hidden(X1,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))|~r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_45, plain, (r2_hidden(X1,k6_enumset1(X2,X3,X4,X5,X6,X7,X8,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_43])])).
cnf(c_0_46, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_45])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 47
# Proof object clause steps            : 20
# Proof object formula steps           : 27
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 19
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 21
# Processed clauses                    : 37
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 37
# Other redundant clauses eliminated   : 17
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 15
# ...of the previous two non-trivial   : 14
# Contextual simplify-reflections      : 0
# Paramodulations                      : 6
# Factorizations                       : 0
# Equation resolutions                 : 17
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 7
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 5
# Current number of unprocessed clauses: 17
# ...number of literals in the above   : 50
# Current number of archived formulas  : 0
# Current number of archived clauses   : 30
# Clause-clause subsumption calls (NU) : 60
# Rec. Clause-clause subsumption calls : 60
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 29
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1657

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.000 s
# Total time               : 0.009 s
# Maximum resident set size: 28800 pages
