# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.006 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t23_tex_2, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&(v7_struct_0(X2)&m1_pre_topc(X2,X1)))=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))=g1_pre_topc(u1_struct_0(k1_tex_2(X1,X3)),u1_pre_topc(k1_tex_2(X1,X3)))))), file('data/MPTP/proofs/26fa430956a34365b81a4493c2ca1e97.E_input', t23_tex_2)).
fof(t65_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>(m1_pre_topc(X1,X2)<=>m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))))), file('data/MPTP/proofs/26fa430956a34365b81a4493c2ca1e97.E_input', t65_pre_topc)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/26fa430956a34365b81a4493c2ca1e97.E_input', dt_m1_pre_topc)).
fof(t5_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X1)=>(u1_struct_0(X2)=u1_struct_0(X3)=>g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))=g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3)))))), file('data/MPTP/proofs/26fa430956a34365b81a4493c2ca1e97.E_input', t5_tsep_1)).
fof(d9_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>(m1_pre_topc(X2,X1)<=>(r1_tarski(k2_struct_0(X2),k2_struct_0(X1))&![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>(r2_hidden(X3,u1_pre_topc(X2))<=>?[X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))&(r2_hidden(X4,u1_pre_topc(X1))&X3=k9_subset_1(u1_struct_0(X2),X4,k2_struct_0(X2)))))))))), file('data/MPTP/proofs/26fa430956a34365b81a4493c2ca1e97.E_input', d9_pre_topc)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/26fa430956a34365b81a4493c2ca1e97.E_input', d3_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/26fa430956a34365b81a4493c2ca1e97.E_input', dt_l1_pre_topc)).
fof(d4_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((~(v2_struct_0(X3))&(v1_pre_topc(X3)&m1_pre_topc(X3,X1)))=>(X3=k1_tex_2(X1,X2)<=>u1_struct_0(X3)=k6_domain_1(u1_struct_0(X1),X2))))), file('data/MPTP/proofs/26fa430956a34365b81a4493c2ca1e97.E_input', d4_tex_2)).
fof(dt_k1_tex_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>(~(v2_struct_0(k1_tex_2(X1,X2)))&(v1_pre_topc(k1_tex_2(X1,X2))&m1_pre_topc(k1_tex_2(X1,X2),X1)))), file('data/MPTP/proofs/26fa430956a34365b81a4493c2ca1e97.E_input', dt_k1_tex_2)).
fof(d1_tex_1, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>(v7_struct_0(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&u1_struct_0(X1)=k6_domain_1(u1_struct_0(X1),X2)))), file('data/MPTP/proofs/26fa430956a34365b81a4493c2ca1e97.E_input', d1_tex_1)).
fof(t59_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))=>m1_pre_topc(X2,X1))), file('data/MPTP/proofs/26fa430956a34365b81a4493c2ca1e97.E_input', t59_pre_topc)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/26fa430956a34365b81a4493c2ca1e97.E_input', d3_tarski)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/26fa430956a34365b81a4493c2ca1e97.E_input', d10_xboole_0)).
fof(t2_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>m1_pre_topc(X1,X1)), file('data/MPTP/proofs/26fa430956a34365b81a4493c2ca1e97.E_input', t2_tsep_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/26fa430956a34365b81a4493c2ca1e97.E_input', t2_subset)).
fof(fc1_struct_0, axiom, ![X1]:((v2_struct_0(X1)&l1_struct_0(X1))=>v1_xboole_0(u1_struct_0(X1))), file('data/MPTP/proofs/26fa430956a34365b81a4493c2ca1e97.E_input', fc1_struct_0)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/26fa430956a34365b81a4493c2ca1e97.E_input', redefinition_k6_domain_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/26fa430956a34365b81a4493c2ca1e97.E_input', d2_subset_1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/26fa430956a34365b81a4493c2ca1e97.E_input', fc2_struct_0)).
fof(c_0_19, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&(v7_struct_0(X2)&m1_pre_topc(X2,X1)))=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))=g1_pre_topc(u1_struct_0(k1_tex_2(X1,X3)),u1_pre_topc(k1_tex_2(X1,X3))))))), inference(assume_negation,[status(cth)],[t23_tex_2])).
fof(c_0_20, plain, ![X8, X9]:((~m1_pre_topc(X8,X9)|m1_pre_topc(X8,g1_pre_topc(u1_struct_0(X9),u1_pre_topc(X9)))|~l1_pre_topc(X9)|~l1_pre_topc(X8))&(~m1_pre_topc(X8,g1_pre_topc(u1_struct_0(X9),u1_pre_topc(X9)))|m1_pre_topc(X8,X9)|~l1_pre_topc(X9)|~l1_pre_topc(X8))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t65_pre_topc])])])])).
fof(c_0_21, plain, ![X10, X11]:(~l1_pre_topc(X10)|(~m1_pre_topc(X11,X10)|l1_pre_topc(X11))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
fof(c_0_22, plain, ![X12, X13, X14]:(~l1_pre_topc(X12)|(~m1_pre_topc(X13,X12)|(~m1_pre_topc(X14,X12)|(u1_struct_0(X13)!=u1_struct_0(X14)|g1_pre_topc(u1_struct_0(X13),u1_pre_topc(X13))=g1_pre_topc(u1_struct_0(X14),u1_pre_topc(X14)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_tsep_1])])])).
fof(c_0_23, negated_conjecture, ![X7]:((~v2_struct_0(esk1_0)&l1_pre_topc(esk1_0))&((~v2_struct_0(esk2_0)&(v7_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0)))&(~m1_subset_1(X7,u1_struct_0(esk1_0))|g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))!=g1_pre_topc(u1_struct_0(k1_tex_2(esk1_0,X7)),u1_pre_topc(k1_tex_2(esk1_0,X7)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_19])])])])])).
fof(c_0_24, plain, ![X15, X16, X17, X19, X21]:(((r1_tarski(k2_struct_0(X16),k2_struct_0(X15))|~m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))&(((m1_subset_1(esk3_3(X15,X16,X17),k1_zfmisc_1(u1_struct_0(X15)))|~r2_hidden(X17,u1_pre_topc(X16))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))&((r2_hidden(esk3_3(X15,X16,X17),u1_pre_topc(X15))|~r2_hidden(X17,u1_pre_topc(X16))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))&(X17=k9_subset_1(u1_struct_0(X16),esk3_3(X15,X16,X17),k2_struct_0(X16))|~r2_hidden(X17,u1_pre_topc(X16))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))))&(~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X15)))|(~r2_hidden(X19,u1_pre_topc(X15))|X17!=k9_subset_1(u1_struct_0(X16),X19,k2_struct_0(X16)))|r2_hidden(X17,u1_pre_topc(X16))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))))&((m1_subset_1(esk4_2(X15,X16),k1_zfmisc_1(u1_struct_0(X16)))|~r1_tarski(k2_struct_0(X16),k2_struct_0(X15))|m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))&((~r2_hidden(esk4_2(X15,X16),u1_pre_topc(X16))|(~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X15)))|(~r2_hidden(X21,u1_pre_topc(X15))|esk4_2(X15,X16)!=k9_subset_1(u1_struct_0(X16),X21,k2_struct_0(X16))))|~r1_tarski(k2_struct_0(X16),k2_struct_0(X15))|m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))&((m1_subset_1(esk5_2(X15,X16),k1_zfmisc_1(u1_struct_0(X15)))|r2_hidden(esk4_2(X15,X16),u1_pre_topc(X16))|~r1_tarski(k2_struct_0(X16),k2_struct_0(X15))|m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))&((r2_hidden(esk5_2(X15,X16),u1_pre_topc(X15))|r2_hidden(esk4_2(X15,X16),u1_pre_topc(X16))|~r1_tarski(k2_struct_0(X16),k2_struct_0(X15))|m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))&(esk4_2(X15,X16)=k9_subset_1(u1_struct_0(X16),esk5_2(X15,X16),k2_struct_0(X16))|r2_hidden(esk4_2(X15,X16),u1_pre_topc(X16))|~r1_tarski(k2_struct_0(X16),k2_struct_0(X15))|m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_pre_topc])])])])])).
fof(c_0_25, plain, ![X23]:(~l1_struct_0(X23)|k2_struct_0(X23)=u1_struct_0(X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
fof(c_0_26, plain, ![X24]:(~l1_pre_topc(X24)|l1_struct_0(X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_27, plain, ![X25, X26, X27]:((X27!=k1_tex_2(X25,X26)|u1_struct_0(X27)=k6_domain_1(u1_struct_0(X25),X26)|(v2_struct_0(X27)|(~v1_pre_topc(X27)|~m1_pre_topc(X27,X25)))|~m1_subset_1(X26,u1_struct_0(X25))|(v2_struct_0(X25)|~l1_pre_topc(X25)))&(u1_struct_0(X27)!=k6_domain_1(u1_struct_0(X25),X26)|X27=k1_tex_2(X25,X26)|(v2_struct_0(X27)|(~v1_pre_topc(X27)|~m1_pre_topc(X27,X25)))|~m1_subset_1(X26,u1_struct_0(X25))|(v2_struct_0(X25)|~l1_pre_topc(X25)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_tex_2])])])])])).
fof(c_0_28, plain, ![X28, X29]:((~v2_struct_0(k1_tex_2(X28,X29))|(v2_struct_0(X28)|~l1_pre_topc(X28)|~m1_subset_1(X29,u1_struct_0(X28))))&((v1_pre_topc(k1_tex_2(X28,X29))|(v2_struct_0(X28)|~l1_pre_topc(X28)|~m1_subset_1(X29,u1_struct_0(X28))))&(m1_pre_topc(k1_tex_2(X28,X29),X28)|(v2_struct_0(X28)|~l1_pre_topc(X28)|~m1_subset_1(X29,u1_struct_0(X28)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_tex_2])])])])).
cnf(c_0_29, plain, (m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, plain, (g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))=g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|u1_struct_0(X2)!=u1_struct_0(X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_33, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_34, plain, (r1_tarski(k2_struct_0(X1),k2_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X1)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_35, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_36, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_37, plain, ![X30, X32]:(((m1_subset_1(esk6_1(X30),u1_struct_0(X30))|~v7_struct_0(X30)|(v2_struct_0(X30)|~l1_struct_0(X30)))&(u1_struct_0(X30)=k6_domain_1(u1_struct_0(X30),esk6_1(X30))|~v7_struct_0(X30)|(v2_struct_0(X30)|~l1_struct_0(X30))))&(~m1_subset_1(X32,u1_struct_0(X30))|u1_struct_0(X30)!=k6_domain_1(u1_struct_0(X30),X32)|v7_struct_0(X30)|(v2_struct_0(X30)|~l1_struct_0(X30)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_tex_1])])])])])])).
cnf(c_0_38, plain, (u1_struct_0(X1)=k6_domain_1(u1_struct_0(X2),X3)|v2_struct_0(X1)|v2_struct_0(X2)|X1!=k1_tex_2(X2,X3)|~v1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_39, plain, (m1_pre_topc(k1_tex_2(X1,X2),X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_40, plain, (v1_pre_topc(k1_tex_2(X1,X2))|v2_struct_0(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_41, plain, (v2_struct_0(X1)|~v2_struct_0(k1_tex_2(X1,X2))|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_42, plain, ![X33, X34]:(~l1_pre_topc(X33)|(~m1_pre_topc(X34,g1_pre_topc(u1_struct_0(X33),u1_pre_topc(X33)))|m1_pre_topc(X34,X33))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t59_pre_topc])])])).
cnf(c_0_43, plain, (m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_44, negated_conjecture, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))|u1_struct_0(X1)!=u1_struct_0(esk2_0)|~m1_pre_topc(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])])).
cnf(c_0_45, negated_conjecture, (l1_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_32]), c_0_33])])).
cnf(c_0_46, plain, (r1_tarski(k2_struct_0(X1),k2_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[c_0_34, c_0_30])).
cnf(c_0_47, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_48, plain, (u1_struct_0(X1)=k6_domain_1(u1_struct_0(X1),esk6_1(X1))|v2_struct_0(X1)|~v7_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_49, plain, (k6_domain_1(u1_struct_0(X1),X2)=u1_struct_0(k1_tex_2(X1,X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_38]), c_0_39]), c_0_40]), c_0_41])).
cnf(c_0_50, plain, (m1_subset_1(esk6_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|~v7_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_51, plain, (m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_52, negated_conjecture, (m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|u1_struct_0(X2)!=u1_struct_0(esk2_0)|~m1_pre_topc(X1,esk2_0)|~m1_pre_topc(X2,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])])).
cnf(c_0_53, plain, (r1_tarski(u1_struct_0(X1),k2_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_30])).
cnf(c_0_54, plain, (u1_struct_0(k1_tex_2(X1,esk6_1(X1)))=u1_struct_0(X1)|v2_struct_0(X1)|~v7_struct_0(X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50]), c_0_36])).
cnf(c_0_55, negated_conjecture, (m1_pre_topc(X1,X2)|u1_struct_0(X2)!=u1_struct_0(esk2_0)|~m1_pre_topc(X1,esk2_0)|~m1_pre_topc(X2,esk1_0)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_56, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_57, plain, (r1_tarski(u1_struct_0(X1),k2_struct_0(X2))|v2_struct_0(X1)|~m1_pre_topc(k1_tex_2(X1,esk6_1(X1)),X2)|~v7_struct_0(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_58, negated_conjecture, (m1_pre_topc(k1_tex_2(esk2_0,X1),X2)|u1_struct_0(X2)!=u1_struct_0(esk2_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_pre_topc(X2,esk1_0)|~l1_pre_topc(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_39]), c_0_45])]), c_0_56])).
cnf(c_0_59, negated_conjecture, (v7_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_60, plain, ![X35, X36, X37, X38, X39]:((~r1_tarski(X35,X36)|(~r2_hidden(X37,X35)|r2_hidden(X37,X36)))&((r2_hidden(esk7_2(X38,X39),X38)|r1_tarski(X38,X39))&(~r2_hidden(esk7_2(X38,X39),X39)|r1_tarski(X38,X39)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_61, negated_conjecture, (r1_tarski(u1_struct_0(esk2_0),k2_struct_0(X1))|u1_struct_0(X1)!=u1_struct_0(esk2_0)|~m1_subset_1(esk6_1(esk2_0),u1_struct_0(esk2_0))|~m1_pre_topc(X1,esk1_0)|~l1_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59]), c_0_45])]), c_0_56])).
cnf(c_0_62, plain, (r1_tarski(k2_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_63, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_64, negated_conjecture, (r1_tarski(u1_struct_0(esk2_0),k2_struct_0(X1))|u1_struct_0(X1)!=u1_struct_0(esk2_0)|~l1_struct_0(esk2_0)|~m1_pre_topc(X1,esk1_0)|~l1_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_50]), c_0_59])]), c_0_56])).
cnf(c_0_65, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_47]), c_0_30])).
fof(c_0_66, plain, ![X41, X42]:(((r1_tarski(X41,X42)|X41!=X42)&(r1_tarski(X42,X41)|X41!=X42))&(~r1_tarski(X41,X42)|~r1_tarski(X42,X41)|X41=X42)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_67, plain, (r1_tarski(X1,X2)|~r2_hidden(esk7_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_68, negated_conjecture, (r2_hidden(X1,k2_struct_0(X2))|u1_struct_0(X2)!=u1_struct_0(esk2_0)|~l1_struct_0(esk2_0)|~r2_hidden(X1,u1_struct_0(esk2_0))|~m1_pre_topc(X2,esk1_0)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_69, plain, (r2_hidden(X1,u1_struct_0(X2))|~r2_hidden(X1,u1_struct_0(X3))|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_63, c_0_65])).
cnf(c_0_70, plain, (r2_hidden(esk7_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_71, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_72, negated_conjecture, (r1_tarski(X1,k2_struct_0(X2))|u1_struct_0(X2)!=u1_struct_0(esk2_0)|~l1_struct_0(esk2_0)|~r2_hidden(esk7_2(X1,k2_struct_0(X2)),u1_struct_0(esk2_0))|~m1_pre_topc(X2,esk1_0)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_73, plain, (r2_hidden(esk7_2(u1_struct_0(X1),X2),u1_struct_0(X3))|r1_tarski(u1_struct_0(X1),X2)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X3)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
fof(c_0_74, plain, ![X43]:(~l1_pre_topc(X43)|m1_pre_topc(X43,X43)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tsep_1])])).
cnf(c_0_75, plain, (u1_struct_0(X1)=k2_struct_0(X2)|~r1_tarski(u1_struct_0(X1),k2_struct_0(X2))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_71, c_0_62])).
cnf(c_0_76, negated_conjecture, (r1_tarski(u1_struct_0(X1),k2_struct_0(X2))|u1_struct_0(X2)!=u1_struct_0(esk2_0)|~l1_struct_0(esk2_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk2_0)|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_45])])).
cnf(c_0_77, plain, (m1_pre_topc(X1,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_78, negated_conjecture, (u1_struct_0(X1)=k2_struct_0(X2)|u1_struct_0(X2)!=u1_struct_0(esk2_0)|~l1_struct_0(esk2_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk2_0)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_30])).
cnf(c_0_79, negated_conjecture, (m1_pre_topc(esk2_0,X1)|u1_struct_0(X1)!=u1_struct_0(esk2_0)|~m1_pre_topc(X1,esk1_0)|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_77]), c_0_45])])).
cnf(c_0_80, negated_conjecture, (u1_struct_0(X1)=k2_struct_0(X2)|u1_struct_0(X2)!=u1_struct_0(esk2_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk2_0)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_36]), c_0_45])])).
cnf(c_0_81, negated_conjecture, (m1_pre_topc(esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_32]), c_0_45])])).
cnf(c_0_82, plain, (k2_struct_0(X1)=k2_struct_0(X2)|~r1_tarski(k2_struct_0(X1),k2_struct_0(X2))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_71, c_0_46])).
cnf(c_0_83, negated_conjecture, (k2_struct_0(X1)=u1_struct_0(esk2_0)|u1_struct_0(X1)!=u1_struct_0(esk2_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_45])])).
cnf(c_0_84, negated_conjecture, (k2_struct_0(X1)=u1_struct_0(esk2_0)|u1_struct_0(X2)!=u1_struct_0(esk2_0)|~r1_tarski(u1_struct_0(esk2_0),k2_struct_0(X1))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X2,esk2_0)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_85, negated_conjecture, (k2_struct_0(X1)=u1_struct_0(esk2_0)|u1_struct_0(X2)!=u1_struct_0(esk2_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X2,esk2_0)|~m1_pre_topc(esk2_0,X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_53]), c_0_30])).
fof(c_0_86, plain, ![X44, X45]:(~m1_subset_1(X44,X45)|(v1_xboole_0(X45)|r2_hidden(X44,X45))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_87, plain, ![X46]:(~v2_struct_0(X46)|~l1_struct_0(X46)|v1_xboole_0(u1_struct_0(X46))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_struct_0])])).
cnf(c_0_88, plain, (u1_struct_0(X1)=k2_struct_0(X2)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_53]), c_0_30])).
cnf(c_0_89, plain, (r2_hidden(X1,u1_struct_0(X2))|~r2_hidden(X1,k2_struct_0(X3))|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_63, c_0_62])).
cnf(c_0_90, negated_conjecture, (k2_struct_0(X1)=u1_struct_0(esk2_0)|~m1_pre_topc(esk2_0,X1)|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_81]), c_0_32]), c_0_45])])).
cnf(c_0_91, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_92, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_87])).
fof(c_0_93, plain, ![X47, X48]:(v1_xboole_0(X47)|~m1_subset_1(X48,X47)|k6_domain_1(X47,X48)=k1_tarski(X48)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_94, negated_conjecture, (k2_struct_0(esk2_0)=u1_struct_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_81]), c_0_81]), c_0_45])])).
cnf(c_0_95, negated_conjecture, (r2_hidden(X1,u1_struct_0(X2))|~r2_hidden(X1,u1_struct_0(esk2_0))|~m1_pre_topc(esk2_0,X3)|~m1_pre_topc(X3,esk2_0)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_96, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(esk6_1(X1),u1_struct_0(X1))|~l1_struct_0(X1)|~v7_struct_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_50]), c_0_92])).
cnf(c_0_97, negated_conjecture, (~m1_subset_1(X1,u1_struct_0(esk1_0))|g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))!=g1_pre_topc(u1_struct_0(k1_tex_2(esk1_0,X1)),u1_pre_topc(k1_tex_2(esk1_0,X1)))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_98, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_99, plain, (v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_struct_0(X1)), inference(spm,[status(thm)],[c_0_92, c_0_36])).
cnf(c_0_100, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk2_0))|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_94]), c_0_45])])).
cnf(c_0_101, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|r2_hidden(esk6_1(esk2_0),u1_struct_0(X1))|~l1_struct_0(esk2_0)|~m1_pre_topc(esk2_0,X2)|~m1_pre_topc(X2,esk2_0)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_59])])).
cnf(c_0_102, negated_conjecture, (u1_struct_0(k1_tex_2(esk1_0,X1))!=u1_struct_0(esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_pre_topc(k1_tex_2(esk1_0,X1),esk1_0)), inference(spm,[status(thm)],[c_0_97, c_0_44])).
cnf(c_0_103, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_104, plain, (u1_struct_0(k1_tex_2(X1,X2))=k1_tarski(X2)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_49]), c_0_99])).
cnf(c_0_105, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,u1_struct_0(X2))|~m1_pre_topc(X2,esk2_0)), inference(spm,[status(thm)],[c_0_63, c_0_100])).
cnf(c_0_106, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|r2_hidden(esk6_1(esk2_0),u1_struct_0(X1))|~l1_struct_0(esk2_0)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_77]), c_0_81]), c_0_45])])).
cnf(c_0_107, plain, (r2_hidden(esk7_2(u1_struct_0(X1),X2),u1_struct_0(X3))|r1_tarski(u1_struct_0(X1),X2)|~m1_pre_topc(X4,X3)|~m1_pre_topc(X1,X4)|~l1_pre_topc(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_73]), c_0_30])).
cnf(c_0_108, negated_conjecture, (u1_struct_0(k1_tex_2(esk1_0,X1))!=u1_struct_0(esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_39]), c_0_33])]), c_0_103])).
cnf(c_0_109, plain, (u1_struct_0(k1_tex_2(X1,X2))=u1_struct_0(k1_tex_2(X3,X2))|v1_xboole_0(u1_struct_0(X3))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X3))|~l1_pre_topc(X1)|~l1_pre_topc(X3)), inference(spm,[status(thm)],[c_0_104, c_0_104])).
fof(c_0_110, plain, ![X49, X50]:(((~m1_subset_1(X50,X49)|r2_hidden(X50,X49)|v1_xboole_0(X49))&(~r2_hidden(X50,X49)|m1_subset_1(X50,X49)|v1_xboole_0(X49)))&((~m1_subset_1(X50,X49)|v1_xboole_0(X50)|~v1_xboole_0(X49))&(~v1_xboole_0(X50)|m1_subset_1(X50,X49)|~v1_xboole_0(X49)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
cnf(c_0_111, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|r2_hidden(esk6_1(esk2_0),u1_struct_0(esk2_0))|~l1_struct_0(esk2_0)|~m1_pre_topc(X1,esk2_0)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_112, negated_conjecture, (r2_hidden(esk7_2(u1_struct_0(X1),X2),u1_struct_0(esk1_0))|r1_tarski(u1_struct_0(X1),X2)|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_32]), c_0_33])])).
cnf(c_0_113, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(X1))|u1_struct_0(k1_tex_2(X1,X2))!=u1_struct_0(esk2_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_33])])).
cnf(c_0_114, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_115, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|r2_hidden(esk6_1(esk2_0),u1_struct_0(esk2_0))|~l1_struct_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_81]), c_0_81]), c_0_45])])).
cnf(c_0_116, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk1_0))|~m1_pre_topc(X1,esk2_0)), inference(spm,[status(thm)],[c_0_67, c_0_112])).
cnf(c_0_117, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(X1))|k1_tarski(X2)!=u1_struct_0(esk2_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_113, c_0_104])).
cnf(c_0_118, plain, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(esk6_1(X1),u1_struct_0(X1))|~l1_struct_0(X1)|~v7_struct_0(X1)), inference(spm,[status(thm)],[c_0_114, c_0_96])).
cnf(c_0_119, plain, (k1_tarski(esk6_1(X1))=u1_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(esk6_1(X1),u1_struct_0(X1))|~v7_struct_0(X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_54]), c_0_99])).
cnf(c_0_120, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|m1_subset_1(esk6_1(esk2_0),u1_struct_0(esk2_0))|~l1_struct_0(esk2_0)), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_121, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,u1_struct_0(X2))|~m1_pre_topc(X2,esk2_0)), inference(spm,[status(thm)],[c_0_63, c_0_116])).
cnf(c_0_122, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(X1))|k1_tarski(esk6_1(X1))!=u1_struct_0(esk2_0)|~m1_subset_1(esk6_1(X1),u1_struct_0(esk1_0))|~v7_struct_0(X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_36])).
cnf(c_0_123, negated_conjecture, (k1_tarski(esk6_1(esk2_0))=u1_struct_0(esk2_0)|v1_xboole_0(u1_struct_0(esk2_0))|~l1_struct_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120]), c_0_59]), c_0_45])])).
fof(c_0_124, plain, ![X51]:(v2_struct_0(X51)|~l1_struct_0(X51)|~v1_xboole_0(u1_struct_0(X51))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_125, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|r2_hidden(esk6_1(esk2_0),u1_struct_0(esk1_0))|~l1_struct_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_115]), c_0_81])])).
cnf(c_0_126, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))|~l1_struct_0(esk2_0)|~m1_subset_1(esk6_1(esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_59]), c_0_45])])).
cnf(c_0_127, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_128, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))|~l1_struct_0(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_125]), c_0_126])).
cnf(c_0_129, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_127, c_0_36])).
cnf(c_0_130, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_36]), c_0_45])])).
cnf(c_0_131, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130]), c_0_45])]), c_0_56])).
cnf(c_0_132, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_131]), c_0_33])]), c_0_103]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 133
# Proof object clause steps            : 94
# Proof object formula steps           : 39
# Proof object conjectures             : 51
# Proof object clause conjectures      : 48
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 29
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 62
# Proof object simplifying inferences  : 84
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 46
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 46
# Processed clauses                    : 16514
# ...of these trivial                  : 640
# ...subsumed                          : 13645
# ...remaining for further processing  : 2229
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 596
# Backward-rewritten                   : 120
# Generated clauses                    : 73596
# ...of the previous two non-trivial   : 68050
# Contextual simplify-reflections      : 489
# Paramodulations                      : 73545
# Factorizations                       : 0
# Equation resolutions                 : 51
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1466
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 1454
# Current number of unprocessed clauses: 48921
# ...number of literals in the above   : 416768
# Current number of archived formulas  : 0
# Current number of archived clauses   : 759
# Clause-clause subsumption calls (NU) : 1055002
# Rec. Clause-clause subsumption calls : 68650
# Non-unit clause-clause subsumptions  : 12194
# Unit Clause-clause subsumption calls : 2334
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2480060

# -------------------------------------------------
# User time                : 1.207 s
# System time              : 0.020 s
# Total time               : 1.227 s
# Maximum resident set size: 49544 pages
