# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_waybel_7, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v1_subset_1(X2,u1_struct_0(X1))<=>~(r2_hidden(k3_yellow_0(X1),X2))))), file('data/MPTP/proofs/272fdf990ce348c9bb1fcfc0f6bfbcdd.E_input', t8_waybel_7)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/272fdf990ce348c9bb1fcfc0f6bfbcdd.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/272fdf990ce348c9bb1fcfc0f6bfbcdd.E_input', d4_subset_1)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/272fdf990ce348c9bb1fcfc0f6bfbcdd.E_input', l3_subset_1)).
fof(t8_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,X2)<=>r2_hidden(X4,X3)))=>X2=X3))), file('data/MPTP/proofs/272fdf990ce348c9bb1fcfc0f6bfbcdd.E_input', t8_subset_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/272fdf990ce348c9bb1fcfc0f6bfbcdd.E_input', d1_xboole_0)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/272fdf990ce348c9bb1fcfc0f6bfbcdd.E_input', dt_k1_yellow_0)).
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/272fdf990ce348c9bb1fcfc0f6bfbcdd.E_input', d9_yellow_0)).
fof(d20_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v13_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r1_orders_2(X1,X3,X4))=>r2_hidden(X4,X2))))))), file('data/MPTP/proofs/272fdf990ce348c9bb1fcfc0f6bfbcdd.E_input', d20_waybel_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/272fdf990ce348c9bb1fcfc0f6bfbcdd.E_input', t4_subset)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/272fdf990ce348c9bb1fcfc0f6bfbcdd.E_input', t2_subset)).
fof(t6_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_lattice3(X1,k1_xboole_0,X2)&r1_lattice3(X1,k1_xboole_0,X2)))), file('data/MPTP/proofs/272fdf990ce348c9bb1fcfc0f6bfbcdd.E_input', t6_yellow_0)).
fof(t42_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))=>(r1_yellow_0(X1,k1_xboole_0)&r2_yellow_0(X1,u1_struct_0(X1)))), file('data/MPTP/proofs/272fdf990ce348c9bb1fcfc0f6bfbcdd.E_input', t42_yellow_0)).
fof(d11_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0)), file('data/MPTP/proofs/272fdf990ce348c9bb1fcfc0f6bfbcdd.E_input', d11_yellow_0)).
fof(d7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(v1_subset_1(X2,X1)<=>~(X2=X1))), file('data/MPTP/proofs/272fdf990ce348c9bb1fcfc0f6bfbcdd.E_input', d7_subset_1)).
fof(fc14_subset_1, axiom, ![X1]:~(v1_subset_1(k2_subset_1(X1),X1)), file('data/MPTP/proofs/272fdf990ce348c9bb1fcfc0f6bfbcdd.E_input', fc14_subset_1)).
fof(c_0_16, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v1_subset_1(X2,u1_struct_0(X1))<=>~(r2_hidden(k3_yellow_0(X1),X2)))))), inference(assume_negation,[status(cth)],[t8_waybel_7])).
fof(c_0_17, plain, ![X30]:m1_subset_1(k2_subset_1(X30),k1_zfmisc_1(X30)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_18, plain, ![X22]:k2_subset_1(X22)=X22, inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_19, plain, ![X32, X33, X34]:(~m1_subset_1(X33,k1_zfmisc_1(X32))|(~r2_hidden(X34,X33)|r2_hidden(X34,X32))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
fof(c_0_20, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_yellow_0(esk1_0)&l1_orders_2(esk1_0))))))&((~v1_xboole_0(esk2_0)&(v2_waybel_0(esk2_0,esk1_0)&(v13_waybel_0(esk2_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))))&((~v1_subset_1(esk2_0,u1_struct_0(esk1_0))|r2_hidden(k3_yellow_0(esk1_0),esk2_0))&(v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k3_yellow_0(esk1_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_16])])])])).
fof(c_0_21, plain, ![X9, X10, X11]:((m1_subset_1(esk3_3(X9,X10,X11),X9)|X10=X11|~m1_subset_1(X11,k1_zfmisc_1(X9))|~m1_subset_1(X10,k1_zfmisc_1(X9)))&((~r2_hidden(esk3_3(X9,X10,X11),X10)|~r2_hidden(esk3_3(X9,X10,X11),X11)|X10=X11|~m1_subset_1(X11,k1_zfmisc_1(X9))|~m1_subset_1(X10,k1_zfmisc_1(X9)))&(r2_hidden(esk3_3(X9,X10,X11),X10)|r2_hidden(esk3_3(X9,X10,X11),X11)|X10=X11|~m1_subset_1(X11,k1_zfmisc_1(X9))|~m1_subset_1(X10,k1_zfmisc_1(X9))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_subset_1])])])])])).
cnf(c_0_22, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_26, plain, ![X26, X27, X28]:((~v1_xboole_0(X26)|~r2_hidden(X27,X26))&(r2_hidden(esk6_1(X28),X28)|v1_xboole_0(X28))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_27, plain, ![X38, X39]:(~l1_orders_2(X38)|m1_subset_1(k1_yellow_0(X38,X39),u1_struct_0(X38))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
cnf(c_0_28, plain, (m1_subset_1(esk3_3(X1,X2,X3),X1)|X2=X3|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_30, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31, plain, (r2_hidden(esk6_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_33, plain, ![X40, X41, X42, X43]:(((r2_lattice3(X40,X41,X42)|X42!=k1_yellow_0(X40,X41)|~r1_yellow_0(X40,X41)|~m1_subset_1(X42,u1_struct_0(X40))|~l1_orders_2(X40))&(~m1_subset_1(X43,u1_struct_0(X40))|(~r2_lattice3(X40,X41,X43)|r1_orders_2(X40,X42,X43))|X42!=k1_yellow_0(X40,X41)|~r1_yellow_0(X40,X41)|~m1_subset_1(X42,u1_struct_0(X40))|~l1_orders_2(X40)))&((m1_subset_1(esk7_3(X40,X41,X42),u1_struct_0(X40))|~r2_lattice3(X40,X41,X42)|X42=k1_yellow_0(X40,X41)|~r1_yellow_0(X40,X41)|~m1_subset_1(X42,u1_struct_0(X40))|~l1_orders_2(X40))&((r2_lattice3(X40,X41,esk7_3(X40,X41,X42))|~r2_lattice3(X40,X41,X42)|X42=k1_yellow_0(X40,X41)|~r1_yellow_0(X40,X41)|~m1_subset_1(X42,u1_struct_0(X40))|~l1_orders_2(X40))&(~r1_orders_2(X40,X42,esk7_3(X40,X41,X42))|~r2_lattice3(X40,X41,X42)|X42=k1_yellow_0(X40,X41)|~r1_yellow_0(X40,X41)|~m1_subset_1(X42,u1_struct_0(X40))|~l1_orders_2(X40))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
cnf(c_0_34, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_36, plain, ![X13, X14, X15, X16]:((~v13_waybel_0(X14,X13)|(~m1_subset_1(X15,u1_struct_0(X13))|(~m1_subset_1(X16,u1_struct_0(X13))|(~r2_hidden(X15,X14)|~r1_orders_2(X13,X15,X16)|r2_hidden(X16,X14))))|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))&((m1_subset_1(esk4_2(X13,X14),u1_struct_0(X13))|v13_waybel_0(X14,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))&((m1_subset_1(esk5_2(X13,X14),u1_struct_0(X13))|v13_waybel_0(X14,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))&(((r2_hidden(esk4_2(X13,X14),X14)|v13_waybel_0(X14,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))&(r1_orders_2(X13,esk4_2(X13,X14),esk5_2(X13,X14))|v13_waybel_0(X14,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13)))&(~r2_hidden(esk5_2(X13,X14),X14)|v13_waybel_0(X14,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d20_waybel_0])])])])])).
fof(c_0_37, plain, ![X35, X36, X37]:(~r2_hidden(X35,X36)|~m1_subset_1(X36,k1_zfmisc_1(X37))|m1_subset_1(X35,X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_38, plain, ![X20, X21]:(~m1_subset_1(X20,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_39, plain, (X1=X2|m1_subset_1(esk3_3(X2,X1,X2),X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_40, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_41, negated_conjecture, (r2_hidden(esk6_1(esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])).
cnf(c_0_42, plain, (r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~r1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_43, negated_conjecture, (m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
fof(c_0_44, plain, ![X7, X8]:((r2_lattice3(X7,k1_xboole_0,X8)|~m1_subset_1(X8,u1_struct_0(X7))|~l1_orders_2(X7))&(r1_lattice3(X7,k1_xboole_0,X8)|~m1_subset_1(X8,u1_struct_0(X7))|~l1_orders_2(X7))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_yellow_0])])])])).
fof(c_0_45, plain, ![X23]:((r1_yellow_0(X23,k1_xboole_0)|(v2_struct_0(X23)|(~v5_orders_2(X23)|(~v1_yellow_0(X23)|~l1_orders_2(X23)))))&(r2_yellow_0(X23,u1_struct_0(X23))|(v2_struct_0(X23)|(~v5_orders_2(X23)|(~v1_yellow_0(X23)|~l1_orders_2(X23)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t42_yellow_0])])])])).
fof(c_0_46, plain, ![X31]:(~l1_orders_2(X31)|k3_yellow_0(X31)=k1_yellow_0(X31,k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_yellow_0])])).
cnf(c_0_47, plain, (r2_hidden(X4,X1)|~v13_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,X1)|~r1_orders_2(X2,X3,X4)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_48, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_49, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|m1_subset_1(esk3_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_39, c_0_25])).
cnf(c_0_51, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_52, negated_conjecture, (r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),X2)|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,X3)|~r1_yellow_0(esk1_0,X3)|~r2_lattice3(esk1_0,X3,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_35])])).
cnf(c_0_53, plain, (r2_lattice3(X1,k1_xboole_0,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_54, plain, (r1_yellow_0(X1,k1_xboole_0)|v2_struct_0(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_55, negated_conjecture, (v1_yellow_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_56, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_57, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_58, plain, (k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_59, plain, ![X24, X25]:((~v1_subset_1(X25,X24)|X25!=X24|~m1_subset_1(X25,k1_zfmisc_1(X24)))&(X25=X24|v1_subset_1(X25,X24)|~m1_subset_1(X25,k1_zfmisc_1(X24)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_subset_1])])])])).
cnf(c_0_60, plain, (r2_hidden(X1,X2)|~r1_orders_2(X3,X4,X1)|~r2_hidden(X4,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~v13_waybel_0(X2,X3)|~l1_orders_2(X3)), inference(csr,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_61, negated_conjecture, (v13_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_62, plain, (X2=X3|~r2_hidden(esk3_3(X1,X2,X3),X2)|~r2_hidden(esk3_3(X1,X2,X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_63, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_hidden(esk3_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51])).
cnf(c_0_64, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),esk3_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)))|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,X2)|~r1_yellow_0(esk1_0,X2)|~r2_lattice3(esk1_0,X2,esk3_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_52, c_0_50])).
cnf(c_0_65, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_lattice3(esk1_0,k1_xboole_0,esk3_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_50]), c_0_35])])).
cnf(c_0_66, negated_conjecture, (r1_yellow_0(esk1_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_35]), c_0_55]), c_0_56])]), c_0_57])).
cnf(c_0_67, negated_conjecture, (r2_hidden(k3_yellow_0(esk1_0),esk2_0)|~v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_68, negated_conjecture, (k3_yellow_0(esk1_0)=k1_yellow_0(esk1_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_58, c_0_35])).
cnf(c_0_69, plain, (X1=X2|v1_subset_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_70, negated_conjecture, (r2_hidden(X1,esk2_0)|~r1_orders_2(esk1_0,X2,X1)|~r2_hidden(X2,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_25]), c_0_61]), c_0_35])])).
cnf(c_0_71, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|~r2_hidden(esk3_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_29]), c_0_25])])).
cnf(c_0_72, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),esk3_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)))|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_66])])).
cnf(c_0_73, negated_conjecture, (r2_hidden(k1_yellow_0(esk1_0,k1_xboole_0),esk2_0)|~v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(rw,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_74, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_69, c_0_25])).
fof(c_0_75, plain, ![X19]:~v1_subset_1(k2_subset_1(X19),X19), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc14_subset_1])])).
cnf(c_0_76, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k3_yellow_0(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_77, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|~r1_orders_2(esk1_0,X1,esk3_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)))|~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_50]), c_0_71])).
cnf(c_0_78, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,k1_xboole_0),esk3_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)))), inference(er,[status(thm)],[c_0_72])).
cnf(c_0_79, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_hidden(k1_yellow_0(esk1_0,k1_xboole_0),esk2_0)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_80, plain, (~v1_subset_1(k2_subset_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_81, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k1_yellow_0(esk1_0,k1_xboole_0),esk2_0)), inference(rw,[status(thm)],[c_0_76, c_0_68])).
cnf(c_0_82, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_79])).
cnf(c_0_83, plain, (~v1_subset_1(X1,X1)), inference(rw,[status(thm)],[c_0_80, c_0_23])).
cnf(c_0_84, negated_conjecture, (r2_hidden(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_43]), c_0_51])).
cnf(c_0_85, negated_conjecture, (~r2_hidden(k1_yellow_0(esk1_0,k1_xboole_0),esk2_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_82]), c_0_83])).
cnf(c_0_86, negated_conjecture, (r2_hidden(k1_yellow_0(esk1_0,X1),esk2_0)), inference(rw,[status(thm)],[c_0_84, c_0_82])).
cnf(c_0_87, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_86])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 88
# Proof object clause steps            : 55
# Proof object formula steps           : 33
# Proof object conjectures             : 37
# Proof object clause conjectures      : 34
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 31
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 42
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 41
# Processed clauses                    : 562
# ...of these trivial                  : 4
# ...subsumed                          : 151
# ...remaining for further processing  : 407
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 217
# Generated clauses                    : 763
# ...of the previous two non-trivial   : 854
# Contextual simplify-reflections      : 7
# Paramodulations                      : 755
# Factorizations                       : 0
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 147
#    Positive orientable unit clauses  : 20
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 124
# Current number of unprocessed clauses: 348
# ...number of literals in the above   : 1685
# Current number of archived formulas  : 0
# Current number of archived clauses   : 260
# Clause-clause subsumption calls (NU) : 16017
# Rec. Clause-clause subsumption calls : 2376
# Non-unit clause-clause subsumptions  : 130
# Unit Clause-clause subsumption calls : 243
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 25
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 27465

# -------------------------------------------------
# User time                : 0.035 s
# System time              : 0.003 s
# Total time               : 0.039 s
# Maximum resident set size: 28740 pages
