# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t116_funct_2, conjecture, ![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X5]:~((r2_hidden(X5,k7_relset_1(X1,X2,X4,X3))&![X6]:(m1_subset_1(X6,X1)=>~((r2_hidden(X6,X3)&X5=k1_funct_1(X4,X6))))))), file('data/MPTP/proofs/275f7d88a24244a48c0062c69bdcbaa5.E_input', t116_funct_2)).
fof(redefinition_k7_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k7_relset_1(X1,X2,X3,X4)=k9_relat_1(X3,X4)), file('data/MPTP/proofs/275f7d88a24244a48c0062c69bdcbaa5.E_input', redefinition_k7_relset_1)).
fof(dt_k7_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k7_relset_1(X1,X2,X3,X4),k1_zfmisc_1(X2))), file('data/MPTP/proofs/275f7d88a24244a48c0062c69bdcbaa5.E_input', dt_k7_relset_1)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/275f7d88a24244a48c0062c69bdcbaa5.E_input', t4_subset)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/275f7d88a24244a48c0062c69bdcbaa5.E_input', d1_xboole_0)).
fof(t143_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k9_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k1_relat_1(X3))&(r2_hidden(k4_tarski(X4,X1),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/275f7d88a24244a48c0062c69bdcbaa5.E_input', t143_relat_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/275f7d88a24244a48c0062c69bdcbaa5.E_input', cc1_relset_1)).
fof(t8_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>(r2_hidden(X1,k1_relat_1(X3))&X2=k1_funct_1(X3,X1)))), file('data/MPTP/proofs/275f7d88a24244a48c0062c69bdcbaa5.E_input', t8_funct_1)).
fof(t52_relset_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(~(v1_xboole_0(X3))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))=>![X5]:(m1_subset_1(X5,X1)=>(r2_hidden(X5,k7_relset_1(X3,X1,X4,X2))<=>?[X6]:(m1_subset_1(X6,X3)&(r2_hidden(k4_tarski(X6,X5),X4)&r2_hidden(X6,X2))))))))), file('data/MPTP/proofs/275f7d88a24244a48c0062c69bdcbaa5.E_input', t52_relset_1)).
fof(cc3_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/275f7d88a24244a48c0062c69bdcbaa5.E_input', cc3_relset_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/275f7d88a24244a48c0062c69bdcbaa5.E_input', t1_subset)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/275f7d88a24244a48c0062c69bdcbaa5.E_input', t2_subset)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/275f7d88a24244a48c0062c69bdcbaa5.E_input', fc1_subset_1)).
fof(cc4_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/275f7d88a24244a48c0062c69bdcbaa5.E_input', cc4_relset_1)).
fof(c_0_14, negated_conjecture, ~(![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X5]:~((r2_hidden(X5,k7_relset_1(X1,X2,X4,X3))&![X6]:(m1_subset_1(X6,X1)=>~((r2_hidden(X6,X3)&X5=k1_funct_1(X4,X6)))))))), inference(assume_negation,[status(cth)],[t116_funct_2])).
fof(c_0_15, plain, ![X13, X14, X15, X16]:(~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14)))|k7_relset_1(X13,X14,X15,X16)=k9_relat_1(X15,X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_relset_1])])).
fof(c_0_16, negated_conjecture, ![X12]:((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,esk1_0,esk2_0)&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))&(r2_hidden(esk5_0,k7_relset_1(esk1_0,esk2_0,esk4_0,esk3_0))&(~m1_subset_1(X12,esk1_0)|(~r2_hidden(X12,esk3_0)|esk5_0!=k1_funct_1(esk4_0,X12))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])])).
fof(c_0_17, plain, ![X17, X18, X19, X20]:(~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(X17,X18)))|m1_subset_1(k7_relset_1(X17,X18,X19,X20),k1_zfmisc_1(X18))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_relset_1])])).
cnf(c_0_18, plain, (k7_relset_1(X2,X3,X1,X4)=k9_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_20, plain, ![X21, X22, X23]:(~r2_hidden(X21,X22)|~m1_subset_1(X22,k1_zfmisc_1(X23))|m1_subset_1(X21,X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_21, plain, ![X24, X25, X26]:((~v1_xboole_0(X24)|~r2_hidden(X25,X24))&(r2_hidden(esk6_1(X26),X26)|v1_xboole_0(X26))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_22, plain, ![X28, X29, X30, X32]:(((r2_hidden(esk7_3(X28,X29,X30),k1_relat_1(X30))|~r2_hidden(X28,k9_relat_1(X30,X29))|~v1_relat_1(X30))&((r2_hidden(k4_tarski(esk7_3(X28,X29,X30),X28),X30)|~r2_hidden(X28,k9_relat_1(X30,X29))|~v1_relat_1(X30))&(r2_hidden(esk7_3(X28,X29,X30),X29)|~r2_hidden(X28,k9_relat_1(X30,X29))|~v1_relat_1(X30))))&(~r2_hidden(X32,k1_relat_1(X30))|(~r2_hidden(k4_tarski(X32,X28),X30)|~r2_hidden(X32,X29))|r2_hidden(X28,k9_relat_1(X30,X29))|~v1_relat_1(X30))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t143_relat_1])])])])])).
fof(c_0_23, plain, ![X33, X34, X35]:(~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(X33,X34)))|v1_relat_1(X35)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_24, plain, (m1_subset_1(k7_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, negated_conjecture, (k7_relset_1(esk1_0,esk2_0,esk4_0,X1)=k9_relat_1(esk4_0,X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
fof(c_0_26, plain, ![X36, X37, X38]:(((r2_hidden(X36,k1_relat_1(X38))|~r2_hidden(k4_tarski(X36,X37),X38)|(~v1_relat_1(X38)|~v1_funct_1(X38)))&(X37=k1_funct_1(X38,X36)|~r2_hidden(k4_tarski(X36,X37),X38)|(~v1_relat_1(X38)|~v1_funct_1(X38))))&(~r2_hidden(X36,k1_relat_1(X38))|X37!=k1_funct_1(X38,X36)|r2_hidden(k4_tarski(X36,X37),X38)|(~v1_relat_1(X38)|~v1_funct_1(X38)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_funct_1])])])).
fof(c_0_27, plain, ![X39, X40, X41, X42, X43, X45]:(((m1_subset_1(esk8_5(X39,X40,X41,X42,X43),X41)|~r2_hidden(X43,k7_relset_1(X41,X39,X42,X40))|~m1_subset_1(X43,X39)|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(X41,X39)))|v1_xboole_0(X41)|v1_xboole_0(X40)|v1_xboole_0(X39))&((r2_hidden(k4_tarski(esk8_5(X39,X40,X41,X42,X43),X43),X42)|~r2_hidden(X43,k7_relset_1(X41,X39,X42,X40))|~m1_subset_1(X43,X39)|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(X41,X39)))|v1_xboole_0(X41)|v1_xboole_0(X40)|v1_xboole_0(X39))&(r2_hidden(esk8_5(X39,X40,X41,X42,X43),X40)|~r2_hidden(X43,k7_relset_1(X41,X39,X42,X40))|~m1_subset_1(X43,X39)|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(X41,X39)))|v1_xboole_0(X41)|v1_xboole_0(X40)|v1_xboole_0(X39))))&(~m1_subset_1(X45,X41)|(~r2_hidden(k4_tarski(X45,X43),X42)|~r2_hidden(X45,X40))|r2_hidden(X43,k7_relset_1(X41,X39,X42,X40))|~m1_subset_1(X43,X39)|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(X41,X39)))|v1_xboole_0(X41)|v1_xboole_0(X40)|v1_xboole_0(X39))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t52_relset_1])])])])])])).
cnf(c_0_28, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_29, plain, ![X46, X47, X48]:(~v1_xboole_0(X46)|(~m1_subset_1(X48,k1_zfmisc_1(k2_zfmisc_1(X46,X47)))|v1_xboole_0(X48))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_relset_1])])])).
fof(c_0_30, plain, ![X49, X50]:(~r2_hidden(X49,X50)|m1_subset_1(X49,X50)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
fof(c_0_31, plain, ![X51, X52]:(~m1_subset_1(X51,X52)|(v1_xboole_0(X52)|r2_hidden(X51,X52))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_32, plain, ![X53]:~v1_xboole_0(k1_zfmisc_1(X53)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
cnf(c_0_33, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_34, plain, (r2_hidden(k4_tarski(esk7_3(X1,X2,X3),X1),X3)|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_35, negated_conjecture, (r2_hidden(esk5_0,k7_relset_1(esk1_0,esk2_0,esk4_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_36, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_37, negated_conjecture, (m1_subset_1(k9_relat_1(esk4_0,X1),k1_zfmisc_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_19])])).
cnf(c_0_38, plain, (X1=k1_funct_1(X2,X3)|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_39, plain, (r2_hidden(k4_tarski(esk8_5(X1,X2,X3,X4,X5),X5),X4)|v1_xboole_0(X3)|v1_xboole_0(X2)|v1_xboole_0(X1)|~r2_hidden(X5,k7_relset_1(X3,X1,X4,X2))|~m1_subset_1(X5,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_40, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relset_1(X3,X2,X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(spm,[status(thm)],[c_0_28, c_0_24])).
cnf(c_0_41, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_42, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_43, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_44, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_45, plain, (~v1_relat_1(X1)|~v1_xboole_0(X1)|~r2_hidden(X2,k9_relat_1(X1,X3))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk5_0,k9_relat_1(esk4_0,esk3_0))), inference(rw,[status(thm)],[c_0_35, c_0_25])).
cnf(c_0_47, negated_conjecture, (v1_relat_1(esk4_0)), inference(spm,[status(thm)],[c_0_36, c_0_19])).
fof(c_0_48, plain, ![X54, X55, X56]:(~v1_xboole_0(X54)|(~m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(X55,X54)))|v1_xboole_0(X56))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])).
cnf(c_0_49, negated_conjecture, (m1_subset_1(X1,esk2_0)|~r2_hidden(X1,k9_relat_1(esk4_0,X2))), inference(spm,[status(thm)],[c_0_28, c_0_37])).
cnf(c_0_50, negated_conjecture, (~m1_subset_1(X1,esk1_0)|~r2_hidden(X1,esk3_0)|esk5_0!=k1_funct_1(esk4_0,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_51, plain, (k1_funct_1(X1,esk8_5(X2,X3,X4,X1,X5))=X5|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X4)|~r2_hidden(X5,k7_relset_1(X4,X2,X1,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40]), c_0_36])).
cnf(c_0_52, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_53, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_54, negated_conjecture, (r2_hidden(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_19]), c_0_44])).
cnf(c_0_55, negated_conjecture, (~v1_xboole_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])])).
cnf(c_0_56, plain, (r2_hidden(esk7_3(X1,X2,X3),X2)|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_57, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_58, negated_conjecture, (m1_subset_1(esk5_0,esk2_0)), inference(spm,[status(thm)],[c_0_49, c_0_46])).
cnf(c_0_59, negated_conjecture, (v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~r2_hidden(esk8_5(X3,X2,X1,esk4_0,esk5_0),esk3_0)|~r2_hidden(esk5_0,k7_relset_1(X1,X3,esk4_0,X2))|~m1_subset_1(esk8_5(X3,X2,X1,esk4_0,esk5_0),esk1_0)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52])])])).
cnf(c_0_60, plain, (m1_subset_1(esk8_5(X1,X2,X3,X4,X5),X3)|v1_xboole_0(X3)|v1_xboole_0(X2)|v1_xboole_0(X1)|~r2_hidden(X5,k7_relset_1(X3,X1,X4,X2))|~m1_subset_1(X5,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_61, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55])).
cnf(c_0_62, plain, (~v1_relat_1(X1)|~v1_xboole_0(X2)|~r2_hidden(X3,k9_relat_1(X1,X2))), inference(spm,[status(thm)],[c_0_33, c_0_56])).
cnf(c_0_63, negated_conjecture, (v1_xboole_0(esk4_0)|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_57, c_0_19])).
cnf(c_0_64, negated_conjecture, (v1_xboole_0(esk2_0)|r2_hidden(esk5_0,esk2_0)), inference(spm,[status(thm)],[c_0_43, c_0_58])).
cnf(c_0_65, negated_conjecture, (v1_xboole_0(X1)|v1_xboole_0(X2)|~r2_hidden(esk8_5(X1,X2,esk1_0,esk4_0,esk5_0),esk3_0)|~r2_hidden(esk5_0,k7_relset_1(esk1_0,X1,esk4_0,X2))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,X1)))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61]), c_0_40])).
cnf(c_0_66, plain, (r2_hidden(esk8_5(X1,X2,X3,X4,X5),X2)|v1_xboole_0(X3)|v1_xboole_0(X2)|v1_xboole_0(X1)|~r2_hidden(X5,k7_relset_1(X3,X1,X4,X2))|~m1_subset_1(X5,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_67, negated_conjecture, (~v1_xboole_0(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_46]), c_0_47])])).
cnf(c_0_68, negated_conjecture, (r2_hidden(esk5_0,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_55])).
cnf(c_0_69, negated_conjecture, (v1_xboole_0(X1)|~r2_hidden(esk5_0,k7_relset_1(esk1_0,X1,esk4_0,esk3_0))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,X1)))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67]), c_0_61]), c_0_40])).
cnf(c_0_70, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_33, c_0_68])).
cnf(c_0_71, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_19]), c_0_25]), c_0_46])]), c_0_70]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 72
# Proof object clause steps            : 43
# Proof object formula steps           : 29
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 22
# Proof object simplifying inferences  : 24
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 3291
# ...of these trivial                  : 2
# ...subsumed                          : 2303
# ...remaining for further processing  : 986
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 32
# Backward-rewritten                   : 6
# Generated clauses                    : 12937
# ...of the previous two non-trivial   : 12816
# Contextual simplify-reflections      : 105
# Paramodulations                      : 12931
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 920
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 894
# Current number of unprocessed clauses: 9526
# ...number of literals in the above   : 74195
# Current number of archived formulas  : 0
# Current number of archived clauses   : 65
# Clause-clause subsumption calls (NU) : 190826
# Rec. Clause-clause subsumption calls : 23733
# Non-unit clause-clause subsumptions  : 2105
# Unit Clause-clause subsumption calls : 329
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 18
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 418902

# -------------------------------------------------
# User time                : 0.360 s
# System time              : 0.003 s
# Total time               : 0.364 s
# Maximum resident set size: 28632 pages
