# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t31_funct_2, conjecture, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>((v2_funct_1(X3)&k2_relset_1(X1,X2,X3)=X2)=>(v1_funct_1(k2_funct_1(X3))&(v1_funct_2(k2_funct_1(X3),X2,X1)&m1_subset_1(k2_funct_1(X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1))))))), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', t31_funct_2)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', t113_zfmisc_1)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', t5_subset)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', fc1_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', d3_tarski)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', d10_xboole_0)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', cc1_relset_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', d1_xboole_0)).
fof(t88_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k7_relat_1(X2,X1),X2)), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', t88_relat_1)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', redefinition_k2_relset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', t3_subset)).
fof(t148_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k2_relat_1(k7_relat_1(X2,X1))=k9_relat_1(X2,X1)), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', t148_relat_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', cc2_relset_1)).
fof(t3_funct_2, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_funct_1(X1)&(v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))))), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', t3_funct_2)).
fof(t147_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r1_tarski(X1,k2_relat_1(X2))=>k9_relat_1(X2,k10_relat_1(X2,X1))=X1)), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', t147_funct_1)).
fof(d18_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(v4_relat_1(X2,X1)<=>r1_tarski(k1_relat_1(X2),X1))), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', d18_relat_1)).
fof(cc3_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', cc3_relset_1)).
fof(t55_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)=>(k2_relat_1(X1)=k1_relat_1(k2_funct_1(X1))&k1_relat_1(X1)=k2_relat_1(k2_funct_1(X1))))), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', t55_funct_1)).
fof(dt_k2_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_relat_1(k2_funct_1(X1))&v1_funct_1(k2_funct_1(X1)))), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', dt_k2_funct_1)).
fof(t8_boole, axiom, ![X1, X2]:~((v1_xboole_0(X1)&(~(X1=X2)&v1_xboole_0(X2)))), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', t8_boole)).
fof(t87_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k1_relat_1(k7_relat_1(X2,X1)),X1)), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', t87_relat_1)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', l13_xboole_0)).
fof(cc1_partfun1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_partfun1(X3,X1))), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', cc1_partfun1)).
fof(cc1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_partfun1(X3,X1))=>(v1_funct_1(X3)&v1_funct_2(X3,X1,X2)))), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', cc1_funct_2)).
fof(t9_funct_2, axiom, ![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(r1_tarski(X2,X3)=>((X2=k1_xboole_0&~(X1=k1_xboole_0))|(v1_funct_1(X4)&(v1_funct_2(X4,X1,X3)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))))))), file('data/MPTP/proofs/2761ea7899c444858d0c414a245832be.E_input', t9_funct_2)).
fof(c_0_25, negated_conjecture, ~(![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>((v2_funct_1(X3)&k2_relset_1(X1,X2,X3)=X2)=>(v1_funct_1(k2_funct_1(X3))&(v1_funct_2(k2_funct_1(X3),X2,X1)&m1_subset_1(k2_funct_1(X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1)))))))), inference(assume_negation,[status(cth)],[t31_funct_2])).
fof(c_0_26, negated_conjecture, ((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,esk1_0,esk2_0)&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))&((v2_funct_1(esk3_0)&k2_relset_1(esk1_0,esk2_0,esk3_0)=esk2_0)&(~v1_funct_1(k2_funct_1(esk3_0))|(~v1_funct_2(k2_funct_1(esk3_0),esk2_0,esk1_0)|~m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])).
fof(c_0_27, plain, ![X8, X9]:((k2_zfmisc_1(X8,X9)!=k1_xboole_0|(X8=k1_xboole_0|X9=k1_xboole_0))&((X8!=k1_xboole_0|k2_zfmisc_1(X8,X9)=k1_xboole_0)&(X9!=k1_xboole_0|k2_zfmisc_1(X8,X9)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_28, plain, ![X10, X11, X12]:(~r2_hidden(X10,X11)|(~m1_subset_1(X11,k1_zfmisc_1(X12))|~v1_xboole_0(X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_30, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_31, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_32, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))|esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_33, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
fof(c_0_34, plain, ![X13, X14, X15, X16, X17]:((~r1_tarski(X13,X14)|(~r2_hidden(X15,X13)|r2_hidden(X15,X14)))&((r2_hidden(esk4_2(X16,X17),X16)|r1_tarski(X16,X17))&(~r2_hidden(esk4_2(X16,X17),X17)|r1_tarski(X16,X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_35, plain, ![X19, X20]:(((r1_tarski(X19,X20)|X19!=X20)&(r1_tarski(X20,X19)|X19!=X20))&(~r1_tarski(X19,X20)|~r1_tarski(X20,X19)|X19=X20)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_36, negated_conjecture, (esk2_0!=k1_xboole_0|~r2_hidden(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])])).
cnf(c_0_37, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_38, plain, ![X21, X22, X23]:(~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(X21,X22)))|v1_relat_1(X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_39, plain, ![X35, X36, X37]:((~v1_xboole_0(X35)|~r2_hidden(X36,X35))&(r2_hidden(esk5_1(X37),X37)|v1_xboole_0(X37))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_40, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41, negated_conjecture, (r1_tarski(esk3_0,X1)|esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_36, c_0_37])).
fof(c_0_42, plain, ![X24, X25]:(~v1_relat_1(X25)|r1_tarski(k7_relat_1(X25,X24),X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t88_relat_1])])).
cnf(c_0_43, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_44, plain, ![X26, X27, X28]:(~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(X26,X27)))|k2_relset_1(X26,X27,X28)=k2_relat_1(X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
fof(c_0_45, plain, ![X42, X43]:((~m1_subset_1(X42,k1_zfmisc_1(X43))|r1_tarski(X42,X43))&(~r1_tarski(X42,X43)|m1_subset_1(X42,k1_zfmisc_1(X43)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_46, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_47, plain, ![X29, X30]:(~v1_relat_1(X30)|k2_relat_1(k7_relat_1(X30,X29))=k9_relat_1(X30,X29)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t148_relat_1])])).
cnf(c_0_48, negated_conjecture, (X1=esk3_0|esk2_0!=k1_xboole_0|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_49, plain, (r1_tarski(k7_relat_1(X1,X2),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_50, negated_conjecture, (v1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_43, c_0_29])).
cnf(c_0_51, negated_conjecture, (k2_relset_1(esk1_0,esk2_0,esk3_0)=esk2_0), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_52, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_53, plain, ![X44, X45, X46]:((v4_relat_1(X46,X44)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(X44,X45))))&(v5_relat_1(X46,X45)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(X44,X45))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_54, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_55, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_46, c_0_37])).
fof(c_0_56, plain, ![X39]:((v1_funct_1(X39)|(~v1_relat_1(X39)|~v1_funct_1(X39)))&((v1_funct_2(X39,k1_relat_1(X39),k2_relat_1(X39))|(~v1_relat_1(X39)|~v1_funct_1(X39)))&(m1_subset_1(X39,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X39),k2_relat_1(X39))))|(~v1_relat_1(X39)|~v1_funct_1(X39))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_funct_2])])])).
cnf(c_0_57, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_58, plain, ![X31, X32]:(~v1_relat_1(X32)|~v1_funct_1(X32)|(~r1_tarski(X31,k2_relat_1(X32))|k9_relat_1(X32,k10_relat_1(X32,X31))=X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t147_funct_1])])).
cnf(c_0_59, plain, (k2_relat_1(k7_relat_1(X1,X2))=k9_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_60, negated_conjecture, (k7_relat_1(esk3_0,X1)=esk3_0|esk2_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])])).
cnf(c_0_61, negated_conjecture, (k2_relat_1(esk3_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_29])])).
fof(c_0_62, plain, ![X52, X53]:((~v4_relat_1(X53,X52)|r1_tarski(k1_relat_1(X53),X52)|~v1_relat_1(X53))&(~r1_tarski(k1_relat_1(X53),X52)|v4_relat_1(X53,X52)|~v1_relat_1(X53))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d18_relat_1])])])).
cnf(c_0_63, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_64, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(X1))|esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_54, c_0_41])).
cnf(c_0_65, plain, (X1=X2|~r1_tarski(X1,X2)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_40, c_0_55])).
cnf(c_0_66, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_67, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_68, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_69, plain, ![X47, X48, X49]:(~v1_xboole_0(X47)|(~m1_subset_1(X49,k1_zfmisc_1(k2_zfmisc_1(X47,X48)))|v1_xboole_0(X49))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_relset_1])])])).
cnf(c_0_70, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_57])).
cnf(c_0_71, plain, (k9_relat_1(X1,k10_relat_1(X1,X2))=X2|~v1_relat_1(X1)|~v1_funct_1(X1)|~r1_tarski(X2,k2_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_72, negated_conjecture, (k9_relat_1(esk3_0,X1)=esk2_0|esk2_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61]), c_0_50])])).
cnf(c_0_73, plain, (r1_tarski(k1_relat_1(X1),X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_74, negated_conjecture, (v4_relat_1(esk3_0,X1)|esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_75, plain, (X1=X2|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_76, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(esk3_0),esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_61]), c_0_50]), c_0_68])])).
cnf(c_0_77, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_78, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_54, c_0_70])).
cnf(c_0_79, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_55, c_0_33])).
cnf(c_0_80, negated_conjecture, (esk2_0=X1|esk2_0!=k1_xboole_0|~r1_tarski(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_50]), c_0_61]), c_0_68])])).
cnf(c_0_81, negated_conjecture, (r1_tarski(k1_relat_1(esk3_0),X1)|esk2_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_50])])).
cnf(c_0_82, negated_conjecture, (k2_zfmisc_1(k1_relat_1(esk3_0),esk2_0)=esk3_0|~v1_xboole_0(k2_zfmisc_1(k1_relat_1(esk3_0),esk2_0))), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_83, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_84, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_40, c_0_79])).
fof(c_0_85, plain, ![X40]:((k2_relat_1(X40)=k1_relat_1(k2_funct_1(X40))|~v2_funct_1(X40)|(~v1_relat_1(X40)|~v1_funct_1(X40)))&(k1_relat_1(X40)=k2_relat_1(k2_funct_1(X40))|~v2_funct_1(X40)|(~v1_relat_1(X40)|~v1_funct_1(X40)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_funct_1])])])).
fof(c_0_86, plain, ![X41]:((v1_relat_1(k2_funct_1(X41))|(~v1_relat_1(X41)|~v1_funct_1(X41)))&(v1_funct_1(k2_funct_1(X41))|(~v1_relat_1(X41)|~v1_funct_1(X41)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_funct_1])])])).
cnf(c_0_87, negated_conjecture, (k1_relat_1(esk3_0)=esk2_0|esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_88, negated_conjecture, (k2_zfmisc_1(k1_relat_1(esk3_0),esk2_0)=esk3_0|~v1_xboole_0(k1_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_89, negated_conjecture, (k1_relat_1(esk3_0)=k1_xboole_0|esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_84, c_0_81])).
cnf(c_0_90, plain, (k2_relat_1(X1)=k1_relat_1(k2_funct_1(X1))|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_91, plain, (v1_funct_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_92, plain, (v1_relat_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_93, negated_conjecture, (r1_tarski(esk2_0,X1)|esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_81, c_0_87])).
cnf(c_0_94, negated_conjecture, (k2_zfmisc_1(k1_xboole_0,esk2_0)=esk3_0|esk2_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_33])])).
fof(c_0_95, plain, ![X50, X51]:(~v1_xboole_0(X50)|(X50=X51|~v1_xboole_0(X51))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_boole])])])).
cnf(c_0_96, plain, (r2_hidden(esk5_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_97, plain, (m1_subset_1(k2_funct_1(X1),k1_zfmisc_1(k2_zfmisc_1(k2_relat_1(X1),k2_relat_1(k2_funct_1(X1)))))|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_90]), c_0_91]), c_0_92])).
cnf(c_0_98, negated_conjecture, (v2_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_99, plain, ![X33, X34]:(~v1_relat_1(X34)|r1_tarski(k1_relat_1(k7_relat_1(X34,X33)),X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t87_relat_1])])).
cnf(c_0_100, negated_conjecture, (esk2_0=X1|esk2_0!=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_65, c_0_93])).
cnf(c_0_101, negated_conjecture, (v1_xboole_0(X1)|esk2_0!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_94]), c_0_33])])).
cnf(c_0_102, plain, (X1=X2|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_103, negated_conjecture, (v1_xboole_0(esk3_0)|esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_36, c_0_96])).
cnf(c_0_104, negated_conjecture, (m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,k2_relat_1(k2_funct_1(esk3_0)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_61]), c_0_50]), c_0_98]), c_0_68])])).
fof(c_0_105, plain, ![X54]:(~v1_xboole_0(X54)|X54=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_106, plain, (r1_tarski(k1_relat_1(k7_relat_1(X1,X2)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_107, negated_conjecture, (esk2_0=X1|esk2_0!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(esk3_0))), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_108, plain, (m1_subset_1(k7_relat_1(X1,X2),k1_zfmisc_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_54, c_0_49])).
cnf(c_0_109, negated_conjecture, (esk3_0=X1|esk2_0!=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_110, negated_conjecture, (v1_xboole_0(k2_funct_1(esk3_0))|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_77, c_0_104])).
cnf(c_0_111, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_112, plain, (~v1_relat_1(X1)|~v1_xboole_0(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))|~r2_hidden(X2,X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_31, c_0_67])).
cnf(c_0_113, plain, (k1_relat_1(k7_relat_1(X1,k1_xboole_0))=k1_xboole_0|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_84, c_0_106])).
cnf(c_0_114, negated_conjecture, (k7_relat_1(esk3_0,X1)=esk2_0|esk2_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_50])])).
cnf(c_0_115, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(X1))|esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_54, c_0_93])).
cnf(c_0_116, negated_conjecture, (k2_funct_1(esk3_0)=esk3_0|~v1_xboole_0(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_111])).
cnf(c_0_117, plain, (~v1_relat_1(X1)|~v1_xboole_0(k1_relat_1(X1))|~r2_hidden(X2,X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_112, c_0_83])).
cnf(c_0_118, negated_conjecture, (k1_relat_1(esk2_0)=k1_xboole_0|esk2_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_50])])).
cnf(c_0_119, negated_conjecture, (v1_relat_1(esk2_0)|esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_43, c_0_115])).
cnf(c_0_120, negated_conjecture, (k2_funct_1(esk3_0)=esk2_0|~v1_xboole_0(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_110]), c_0_111])).
cnf(c_0_121, plain, (k1_relat_1(X1)=k2_relat_1(k2_funct_1(X1))|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_122, negated_conjecture, (k1_relat_1(esk3_0)=esk2_0|~v1_xboole_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_116]), c_0_61]), c_0_50]), c_0_98]), c_0_68])])).
cnf(c_0_123, negated_conjecture, (esk2_0!=k1_xboole_0|~r2_hidden(X1,esk2_0)|~v1_funct_1(esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_33])]), c_0_119])).
cnf(c_0_124, negated_conjecture, (v1_funct_1(esk2_0)|~v1_xboole_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_120]), c_0_50]), c_0_68])])).
cnf(c_0_125, negated_conjecture, (m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_relat_1(esk3_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_121]), c_0_50]), c_0_98]), c_0_68])])).
cnf(c_0_126, negated_conjecture, (k2_zfmisc_1(esk2_0,esk2_0)=esk3_0|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_88, c_0_122])).
cnf(c_0_127, negated_conjecture, (v1_xboole_0(esk2_0)|esk2_0!=k1_xboole_0|~v1_funct_1(esk2_0)), inference(spm,[status(thm)],[c_0_123, c_0_96])).
cnf(c_0_128, negated_conjecture, (v1_funct_1(esk2_0)|esk2_0!=k1_xboole_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_101]), c_0_115])).
cnf(c_0_129, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1)), inference(spm,[status(thm)],[c_0_63, c_0_78])).
cnf(c_0_130, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_43, c_0_78])).
cnf(c_0_131, negated_conjecture, (m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk2_0)))|esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_125, c_0_87])).
cnf(c_0_132, negated_conjecture, (k2_zfmisc_1(esk2_0,esk2_0)=esk3_0|esk2_0!=k1_xboole_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_127]), c_0_128])).
cnf(c_0_133, plain, (r1_tarski(k1_relat_1(k2_zfmisc_1(X1,X2)),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_129]), c_0_130])])).
fof(c_0_134, plain, ![X58, X59, X60]:(~v1_xboole_0(X58)|(~m1_subset_1(X60,k1_zfmisc_1(k2_zfmisc_1(X58,X59)))|v1_partfun1(X60,X58))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_partfun1])])])).
cnf(c_0_135, negated_conjecture, (m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(esk3_0))|esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_131, c_0_132])).
cnf(c_0_136, plain, (k1_relat_1(k2_zfmisc_1(k1_xboole_0,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_84, c_0_133])).
cnf(c_0_137, negated_conjecture, (~v1_xboole_0(k2_zfmisc_1(esk2_0,k2_relat_1(k2_funct_1(esk3_0))))|~r2_hidden(X1,k2_funct_1(esk3_0))), inference(spm,[status(thm)],[c_0_31, c_0_104])).
fof(c_0_138, plain, ![X55, X56, X57]:((v1_funct_1(X57)|(~v1_funct_1(X57)|~v1_partfun1(X57,X55))|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(X55,X56))))&(v1_funct_2(X57,X55,X56)|(~v1_funct_1(X57)|~v1_partfun1(X57,X55))|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(X55,X56))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_2])])])).
cnf(c_0_139, plain, (v1_partfun1(X2,X1)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_140, negated_conjecture, (m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk2_0)))|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_125, c_0_122])).
cnf(c_0_141, negated_conjecture, (esk2_0!=k1_xboole_0|~r2_hidden(X1,k2_funct_1(esk3_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_135]), c_0_103])).
cnf(c_0_142, negated_conjecture, (k2_funct_1(esk3_0)=esk2_0|esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_107, c_0_135])).
cnf(c_0_143, plain, (k1_relat_1(X1)=k1_xboole_0|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_136, c_0_102])).
cnf(c_0_144, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_145, negated_conjecture, (k2_relat_1(k2_funct_1(esk3_0))!=k1_xboole_0|~r2_hidden(X1,k2_funct_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_30]), c_0_33])])).
cnf(c_0_146, plain, (v1_funct_2(X1,X2,X3)|~v1_funct_1(X1)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_138])).
cnf(c_0_147, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_148, negated_conjecture, (v1_partfun1(k2_funct_1(esk3_0),esk2_0)|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_139, c_0_104])).
cnf(c_0_149, negated_conjecture, (k2_funct_1(esk3_0)=k2_zfmisc_1(esk2_0,esk2_0)|~v1_xboole_0(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_140]), c_0_83])).
cnf(c_0_150, negated_conjecture, (esk2_0!=k1_xboole_0|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_141, c_0_142])).
cnf(c_0_151, plain, (v1_relat_1(k1_xboole_0)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_130, c_0_30])).
cnf(c_0_152, plain, (k1_relat_1(X1)=k1_xboole_0|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144]), c_0_33])])).
cnf(c_0_153, negated_conjecture, (v1_xboole_0(k2_funct_1(esk3_0))|k2_relat_1(k2_funct_1(esk3_0))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_145, c_0_96])).
cnf(c_0_154, negated_conjecture, (~v1_funct_1(k2_funct_1(esk3_0))|~v1_funct_2(k2_funct_1(esk3_0),esk2_0,esk1_0)|~m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_155, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~v1_xboole_0(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_146, c_0_147])).
cnf(c_0_156, negated_conjecture, (v1_partfun1(esk2_0,esk2_0)|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_148, c_0_120])).
cnf(c_0_157, negated_conjecture, (k2_zfmisc_1(esk2_0,esk2_0)=esk2_0|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_120, c_0_149])).
cnf(c_0_158, negated_conjecture, (v1_xboole_0(esk2_0)|esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_150, c_0_96])).
cnf(c_0_159, plain, (v1_relat_1(k1_xboole_0)), inference(er,[status(thm)],[c_0_151])).
cnf(c_0_160, plain, (k2_relat_1(X1)=k1_xboole_0|~v1_relat_1(X1)|~v1_xboole_0(k2_funct_1(X1))|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_90, c_0_152])).
cnf(c_0_161, negated_conjecture, (v1_xboole_0(k2_funct_1(esk3_0))|k1_relat_1(esk3_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_121]), c_0_50]), c_0_98]), c_0_68])])).
cnf(c_0_162, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_111, c_0_83])).
cnf(c_0_163, negated_conjecture, (~v1_xboole_0(esk2_0)|~v1_funct_2(esk2_0,esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_120]), c_0_124]), c_0_147])).
cnf(c_0_164, negated_conjecture, (v1_funct_2(esk2_0,esk2_0,X1)|~v1_xboole_0(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_156]), c_0_124])).
cnf(c_0_165, negated_conjecture, (k2_zfmisc_1(esk2_0,esk2_0)=esk2_0|esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_157, c_0_158])).
cnf(c_0_166, plain, (r1_tarski(k1_relat_1(X1),X2)|~v1_relat_1(X3)|~v1_xboole_0(k7_relat_1(X3,X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_106, c_0_102])).
cnf(c_0_167, plain, (k7_relat_1(k1_xboole_0,X1)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_49]), c_0_159])])).
cnf(c_0_168, plain, (v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
fof(c_0_169, plain, ![X61, X62, X63, X64]:(((v1_funct_1(X64)|X62=k1_xboole_0|~r1_tarski(X62,X63)|(~v1_funct_1(X64)|(~v1_funct_2(X64,X61,X62)|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(X61,X62))))))&((v1_funct_2(X64,X61,X63)|X62=k1_xboole_0|~r1_tarski(X62,X63)|(~v1_funct_1(X64)|(~v1_funct_2(X64,X61,X62)|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(X61,X62))))))&(m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(X61,X63)))|X62=k1_xboole_0|~r1_tarski(X62,X63)|(~v1_funct_1(X64)|(~v1_funct_2(X64,X61,X62)|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(X61,X62))))))))&((v1_funct_1(X64)|X61!=k1_xboole_0|~r1_tarski(X62,X63)|(~v1_funct_1(X64)|(~v1_funct_2(X64,X61,X62)|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(X61,X62))))))&((v1_funct_2(X64,X61,X63)|X61!=k1_xboole_0|~r1_tarski(X62,X63)|(~v1_funct_1(X64)|(~v1_funct_2(X64,X61,X62)|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(X61,X62))))))&(m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(X61,X63)))|X61!=k1_xboole_0|~r1_tarski(X62,X63)|(~v1_funct_1(X64)|(~v1_funct_2(X64,X61,X62)|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(X61,X62))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_funct_2])])])])).
cnf(c_0_170, negated_conjecture, (esk2_0=k1_xboole_0|k1_relat_1(esk3_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_161]), c_0_61]), c_0_50]), c_0_98]), c_0_68])])).
cnf(c_0_171, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_162, c_0_33])).
cnf(c_0_172, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_163, c_0_164])).
cnf(c_0_173, negated_conjecture, (v1_xboole_0(X1)|esk2_0!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_165]), c_0_158])).
cnf(c_0_174, plain, (r1_tarski(k1_relat_1(X1),X2)|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_167]), c_0_159]), c_0_33])])).
cnf(c_0_175, plain, (v1_funct_2(k2_funct_1(X1),k2_relat_1(X1),k2_relat_1(k2_funct_1(X1)))|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_90]), c_0_91]), c_0_92])).
cnf(c_0_176, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|X4=k1_xboole_0|~r1_tarski(X4,X3)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))), inference(split_conjunct,[status(thm)],[c_0_169])).
cnf(c_0_177, negated_conjecture, (v1_funct_2(esk3_0,k1_relat_1(esk3_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_61]), c_0_50]), c_0_68])])).
cnf(c_0_178, negated_conjecture, (esk2_0=k1_xboole_0|~v1_xboole_0(esk3_0)), inference(spm,[status(thm)],[c_0_170, c_0_152])).
cnf(c_0_179, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_171]), c_0_33])])).
cnf(c_0_180, negated_conjecture, (esk2_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_173]), c_0_78])])).
cnf(c_0_181, plain, (r1_tarski(k2_relat_1(X1),X2)|~v1_relat_1(X1)|~v1_xboole_0(k2_funct_1(X1))|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_174, c_0_90])).
cnf(c_0_182, negated_conjecture, (k2_funct_1(esk3_0)=k1_xboole_0|k1_relat_1(esk3_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_111, c_0_161])).
cnf(c_0_183, negated_conjecture, (v1_funct_2(k2_funct_1(esk3_0),esk2_0,k2_relat_1(k2_funct_1(esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_61]), c_0_50]), c_0_98]), c_0_68])])).
cnf(c_0_184, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(esk3_0),X1)))|~r1_tarski(esk2_0,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_76]), c_0_177]), c_0_68])]), c_0_64])).
cnf(c_0_185, negated_conjecture, (~m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_179]), c_0_180])).
cnf(c_0_186, negated_conjecture, (r1_tarski(esk2_0,X1)|k1_relat_1(esk3_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_182]), c_0_61]), c_0_50]), c_0_33]), c_0_98]), c_0_68])])).
cnf(c_0_187, plain, (v1_funct_2(X1,X2,X3)|X4=k1_xboole_0|~r1_tarski(X4,X3)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))), inference(split_conjunct,[status(thm)],[c_0_169])).
cnf(c_0_188, negated_conjecture, (v1_funct_2(k2_funct_1(esk3_0),esk2_0,k1_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_121]), c_0_50]), c_0_98]), c_0_68])])).
cnf(c_0_189, negated_conjecture, (k1_relat_1(esk3_0)!=k1_xboole_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_144]), c_0_185]), c_0_186])).
cnf(c_0_190, negated_conjecture, (v4_relat_1(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_63, c_0_29])).
cnf(c_0_191, negated_conjecture, (v1_funct_2(k2_funct_1(esk3_0),esk2_0,X1)|~r1_tarski(k1_relat_1(esk3_0),X1)|~v1_funct_1(k2_funct_1(esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_125]), c_0_188])]), c_0_189])).
cnf(c_0_192, negated_conjecture, (r1_tarski(k1_relat_1(esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_190]), c_0_50])])).
cnf(c_0_193, negated_conjecture, (~m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))|~v1_funct_1(k2_funct_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_191]), c_0_192])])).
cnf(c_0_194, negated_conjecture, (m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,X1)))|~r1_tarski(k1_relat_1(esk3_0),X1)|~v1_funct_1(k2_funct_1(esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_125]), c_0_188])]), c_0_189])).
cnf(c_0_195, negated_conjecture, (~v1_funct_1(k2_funct_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_194]), c_0_192])])).
cnf(c_0_196, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_91]), c_0_50]), c_0_68])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 197
# Proof object clause steps            : 147
# Proof object formula steps           : 50
# Proof object conjectures             : 86
# Proof object clause conjectures      : 83
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 37
# Proof object initial formulas used   : 25
# Proof object generating inferences   : 109
# Proof object simplifying inferences  : 123
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 25
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 50
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 46
# Processed clauses                    : 21645
# ...of these trivial                  : 34
# ...subsumed                          : 19387
# ...remaining for further processing  : 2224
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 285
# Backward-rewritten                   : 14
# Generated clauses                    : 63729
# ...of the previous two non-trivial   : 58567
# Contextual simplify-reflections      : 419
# Paramodulations                      : 63701
# Factorizations                       : 0
# Equation resolutions                 : 28
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1878
#    Positive orientable unit clauses  : 45
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 14
#    Non-unit-clauses                  : 1819
# Current number of unprocessed clauses: 35905
# ...number of literals in the above   : 175923
# Current number of archived formulas  : 0
# Current number of archived clauses   : 344
# Clause-clause subsumption calls (NU) : 760445
# Rec. Clause-clause subsumption calls : 450289
# Non-unit clause-clause subsumptions  : 13925
# Unit Clause-clause subsumption calls : 5023
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 30
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 819751

# -------------------------------------------------
# User time                : 0.806 s
# System time              : 0.013 s
# Total time               : 0.819 s
# Maximum resident set size: 33964 pages
