# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>(r2_hidden(X1,k1_relat_1(X3))&X2=k1_funct_1(X3,X1)))), file('data/MPTP/proofs/27cfc50634d84e0e99934e613e5f45fc.E_input', t8_funct_1)).
fof(d8_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(X3=k5_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>?[X6]:(r2_hidden(k4_tarski(X4,X6),X1)&r2_hidden(k4_tarski(X6,X5),X2))))))), file('data/MPTP/proofs/27cfc50634d84e0e99934e613e5f45fc.E_input', d8_relat_1)).
fof(d4_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2, X3]:((r2_hidden(X2,k1_relat_1(X1))=>(X3=k1_funct_1(X1,X2)<=>r2_hidden(k4_tarski(X2,X3),X1)))&(~(r2_hidden(X2,k1_relat_1(X1)))=>(X3=k1_funct_1(X1,X2)<=>X3=k1_xboole_0)))), file('data/MPTP/proofs/27cfc50634d84e0e99934e613e5f45fc.E_input', d4_funct_1)).
fof(dt_k5_relat_1, axiom, ![X1, X2]:((v1_relat_1(X1)&v1_relat_1(X2))=>v1_relat_1(k5_relat_1(X1,X2))), file('data/MPTP/proofs/27cfc50634d84e0e99934e613e5f45fc.E_input', dt_k5_relat_1)).
fof(fc2_funct_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&(v1_relat_1(X2)&v1_funct_1(X2)))=>(v1_relat_1(k5_relat_1(X1,X2))&v1_funct_1(k5_relat_1(X1,X2)))), file('data/MPTP/proofs/27cfc50634d84e0e99934e613e5f45fc.E_input', fc2_funct_1)).
fof(t21_funct_1, conjecture, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k1_relat_1(k5_relat_1(X3,X2)))<=>(r2_hidden(X1,k1_relat_1(X3))&r2_hidden(k1_funct_1(X3,X1),k1_relat_1(X2)))))), file('data/MPTP/proofs/27cfc50634d84e0e99934e613e5f45fc.E_input', t21_funct_1)).
fof(c_0_6, plain, ![X12, X13, X14]:(((r2_hidden(X12,k1_relat_1(X14))|~r2_hidden(k4_tarski(X12,X13),X14)|(~v1_relat_1(X14)|~v1_funct_1(X14)))&(X13=k1_funct_1(X14,X12)|~r2_hidden(k4_tarski(X12,X13),X14)|(~v1_relat_1(X14)|~v1_funct_1(X14))))&(~r2_hidden(X12,k1_relat_1(X14))|X13!=k1_funct_1(X14,X12)|r2_hidden(k4_tarski(X12,X13),X14)|(~v1_relat_1(X14)|~v1_funct_1(X14)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_funct_1])])])).
fof(c_0_7, plain, ![X15, X16, X17, X18, X19, X21, X22, X23, X26]:((((r2_hidden(k4_tarski(X18,esk4_5(X15,X16,X17,X18,X19)),X15)|~r2_hidden(k4_tarski(X18,X19),X17)|X17!=k5_relat_1(X15,X16)|~v1_relat_1(X17)|~v1_relat_1(X16)|~v1_relat_1(X15))&(r2_hidden(k4_tarski(esk4_5(X15,X16,X17,X18,X19),X19),X16)|~r2_hidden(k4_tarski(X18,X19),X17)|X17!=k5_relat_1(X15,X16)|~v1_relat_1(X17)|~v1_relat_1(X16)|~v1_relat_1(X15)))&(~r2_hidden(k4_tarski(X21,X23),X15)|~r2_hidden(k4_tarski(X23,X22),X16)|r2_hidden(k4_tarski(X21,X22),X17)|X17!=k5_relat_1(X15,X16)|~v1_relat_1(X17)|~v1_relat_1(X16)|~v1_relat_1(X15)))&((~r2_hidden(k4_tarski(esk5_3(X15,X16,X17),esk6_3(X15,X16,X17)),X17)|(~r2_hidden(k4_tarski(esk5_3(X15,X16,X17),X26),X15)|~r2_hidden(k4_tarski(X26,esk6_3(X15,X16,X17)),X16))|X17=k5_relat_1(X15,X16)|~v1_relat_1(X17)|~v1_relat_1(X16)|~v1_relat_1(X15))&((r2_hidden(k4_tarski(esk5_3(X15,X16,X17),esk7_3(X15,X16,X17)),X15)|r2_hidden(k4_tarski(esk5_3(X15,X16,X17),esk6_3(X15,X16,X17)),X17)|X17=k5_relat_1(X15,X16)|~v1_relat_1(X17)|~v1_relat_1(X16)|~v1_relat_1(X15))&(r2_hidden(k4_tarski(esk7_3(X15,X16,X17),esk6_3(X15,X16,X17)),X16)|r2_hidden(k4_tarski(esk5_3(X15,X16,X17),esk6_3(X15,X16,X17)),X17)|X17=k5_relat_1(X15,X16)|~v1_relat_1(X17)|~v1_relat_1(X16)|~v1_relat_1(X15))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_relat_1])])])])])])).
cnf(c_0_8, plain, (X1=k1_funct_1(X2,X3)|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, plain, (r2_hidden(k4_tarski(esk4_5(X1,X2,X3,X4,X5),X5),X2)|~r2_hidden(k4_tarski(X4,X5),X3)|X3!=k5_relat_1(X1,X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10, plain, (r2_hidden(k4_tarski(X1,esk4_5(X2,X3,X4,X1,X5)),X2)|~r2_hidden(k4_tarski(X1,X5),X4)|X4!=k5_relat_1(X2,X3)|~v1_relat_1(X4)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_11, plain, ![X30, X31, X32]:(((X32!=k1_funct_1(X30,X31)|r2_hidden(k4_tarski(X31,X32),X30)|~r2_hidden(X31,k1_relat_1(X30))|(~v1_relat_1(X30)|~v1_funct_1(X30)))&(~r2_hidden(k4_tarski(X31,X32),X30)|X32=k1_funct_1(X30,X31)|~r2_hidden(X31,k1_relat_1(X30))|(~v1_relat_1(X30)|~v1_funct_1(X30))))&((X32!=k1_funct_1(X30,X31)|X32=k1_xboole_0|r2_hidden(X31,k1_relat_1(X30))|(~v1_relat_1(X30)|~v1_funct_1(X30)))&(X32!=k1_xboole_0|X32=k1_funct_1(X30,X31)|r2_hidden(X31,k1_relat_1(X30))|(~v1_relat_1(X30)|~v1_funct_1(X30))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_funct_1])])])])])).
cnf(c_0_12, plain, (k1_funct_1(X1,esk4_5(X2,X1,X3,X4,X5))=X5|X3!=k5_relat_1(X2,X1)|~r2_hidden(k4_tarski(X4,X5),X3)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_13, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X1,X4)|X3!=k5_relat_1(X1,X2)|~r2_hidden(k4_tarski(X4,X5),X3)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_8, c_0_10])).
cnf(c_0_14, plain, (r2_hidden(k4_tarski(X3,X1),X2)|X1!=k1_funct_1(X2,X3)|~r2_hidden(X3,k1_relat_1(X2))|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (k1_funct_1(X1,k1_funct_1(X2,X3))=X4|X5!=k5_relat_1(X2,X1)|~r2_hidden(k4_tarski(X3,X4),X5)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X5)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_16, plain, (r2_hidden(k4_tarski(X1,k1_funct_1(X2,X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_14])).
fof(c_0_17, plain, ![X10, X11]:(~v1_relat_1(X10)|~v1_relat_1(X11)|v1_relat_1(k5_relat_1(X10,X11))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_relat_1])])).
fof(c_0_18, plain, ![X28, X29]:((v1_relat_1(k5_relat_1(X28,X29))|(~v1_relat_1(X28)|~v1_funct_1(X28)|(~v1_relat_1(X29)|~v1_funct_1(X29))))&(v1_funct_1(k5_relat_1(X28,X29))|(~v1_relat_1(X28)|~v1_funct_1(X28)|(~v1_relat_1(X29)|~v1_funct_1(X29))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc2_funct_1])])])).
cnf(c_0_19, plain, (k1_funct_1(X1,k1_funct_1(X2,X3))=k1_funct_1(X4,X3)|X4!=k5_relat_1(X2,X1)|~r2_hidden(X3,k1_relat_1(X4))|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X4)|~v1_relat_1(X1)|~v1_relat_1(X4)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_20, plain, (v1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21, plain, (v1_funct_1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, plain, (k1_funct_1(k5_relat_1(X1,X2),X3)=k1_funct_1(X2,k1_funct_1(X1,X3))|~r2_hidden(X3,k1_relat_1(k5_relat_1(X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_19]), c_0_20]), c_0_21])).
cnf(c_0_23, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_24, plain, (r2_hidden(k4_tarski(X1,X2),k5_relat_1(X3,X4))|X2!=k1_funct_1(X4,k1_funct_1(X3,X1))|~r2_hidden(X1,k1_relat_1(k5_relat_1(X3,X4)))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_relat_1(X4)|~v1_relat_1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_22]), c_0_20]), c_0_21])).
fof(c_0_25, negated_conjecture, ~(![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k1_relat_1(k5_relat_1(X3,X2)))<=>(r2_hidden(X1,k1_relat_1(X3))&r2_hidden(k1_funct_1(X3,X1),k1_relat_1(X2))))))), inference(assume_negation,[status(cth)],[t21_funct_1])).
cnf(c_0_26, plain, (r2_hidden(k4_tarski(X1,X4),X6)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(k4_tarski(X2,X4),X5)|X6!=k5_relat_1(X3,X5)|~v1_relat_1(X6)|~v1_relat_1(X5)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_27, plain, (r2_hidden(X1,k1_relat_1(X2))|X3!=k5_relat_1(X2,X4)|~r2_hidden(k4_tarski(X1,X5),X3)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X4)), inference(spm,[status(thm)],[c_0_23, c_0_10])).
cnf(c_0_28, plain, (r2_hidden(k4_tarski(X1,k1_funct_1(X2,k1_funct_1(X3,X1))),k5_relat_1(X3,X2))|~r2_hidden(X1,k1_relat_1(k5_relat_1(X3,X2)))|~v1_funct_1(X2)|~v1_funct_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X3)), inference(er,[status(thm)],[c_0_24])).
fof(c_0_29, negated_conjecture, ((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&((v1_relat_1(esk3_0)&v1_funct_1(esk3_0))&((~r2_hidden(esk1_0,k1_relat_1(k5_relat_1(esk3_0,esk2_0)))|(~r2_hidden(esk1_0,k1_relat_1(esk3_0))|~r2_hidden(k1_funct_1(esk3_0,esk1_0),k1_relat_1(esk2_0))))&((r2_hidden(esk1_0,k1_relat_1(esk3_0))|r2_hidden(esk1_0,k1_relat_1(k5_relat_1(esk3_0,esk2_0))))&(r2_hidden(k1_funct_1(esk3_0,esk1_0),k1_relat_1(esk2_0))|r2_hidden(esk1_0,k1_relat_1(k5_relat_1(esk3_0,esk2_0)))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])])).
cnf(c_0_30, plain, (r2_hidden(k4_tarski(X1,k1_funct_1(X2,X3)),X4)|X4!=k5_relat_1(X5,X2)|~r2_hidden(k4_tarski(X1,X3),X5)|~r2_hidden(X3,k1_relat_1(X2))|~v1_funct_1(X2)|~v1_relat_1(X4)|~v1_relat_1(X2)|~v1_relat_1(X5)), inference(spm,[status(thm)],[c_0_26, c_0_16])).
cnf(c_0_31, plain, (r2_hidden(X1,k1_relat_1(X2))|k5_relat_1(X3,X4)!=k5_relat_1(X2,X5)|~r2_hidden(X1,k1_relat_1(k5_relat_1(X3,X4)))|~v1_funct_1(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X5)|~v1_relat_1(X4)|~v1_relat_1(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_20])).
cnf(c_0_32, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(esk3_0))|r2_hidden(esk1_0,k1_relat_1(k5_relat_1(esk3_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_33, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_35, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, plain, (r2_hidden(k4_tarski(X1,k1_funct_1(X2,k1_funct_1(X3,X1))),X4)|X4!=k5_relat_1(X3,X2)|~r2_hidden(k1_funct_1(X3,X1),k1_relat_1(X2))|~r2_hidden(X1,k1_relat_1(X3))|~v1_funct_1(X2)|~v1_funct_1(X3)|~v1_relat_1(X4)|~v1_relat_1(X2)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_30, c_0_16])).
cnf(c_0_38, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(esk3_0))|r2_hidden(esk1_0,k1_relat_1(X1))|k5_relat_1(esk3_0,esk2_0)!=k5_relat_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33]), c_0_34]), c_0_35]), c_0_36])])).
cnf(c_0_39, plain, (r2_hidden(esk4_5(X1,X2,X3,X4,X5),k1_relat_1(X2))|X3!=k5_relat_1(X1,X2)|~r2_hidden(k4_tarski(X4,X5),X3)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_23, c_0_9])).
cnf(c_0_40, plain, (r2_hidden(X1,k1_relat_1(X2))|X2!=k5_relat_1(X3,X4)|~r2_hidden(k1_funct_1(X3,X1),k1_relat_1(X4))|~r2_hidden(X1,k1_relat_1(X3))|~v1_funct_1(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X4)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_23, c_0_37])).
cnf(c_0_41, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk1_0),k1_relat_1(esk2_0))|r2_hidden(esk1_0,k1_relat_1(k5_relat_1(esk3_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_42, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_38]), c_0_34]), c_0_36]), c_0_35])])).
cnf(c_0_43, plain, (r2_hidden(k1_funct_1(X1,X2),k1_relat_1(X3))|X4!=k5_relat_1(X1,X3)|~r2_hidden(k4_tarski(X2,X5),X4)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X4)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_39, c_0_13])).
cnf(c_0_44, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(k5_relat_1(esk3_0,esk2_0)))|r2_hidden(esk1_0,k1_relat_1(X1))|X1!=k5_relat_1(esk3_0,esk2_0)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42]), c_0_33]), c_0_34]), c_0_35]), c_0_36])])).
cnf(c_0_45, plain, (r2_hidden(k1_funct_1(X1,X2),k1_relat_1(X3))|X4!=k5_relat_1(X1,X3)|~r2_hidden(X2,k1_relat_1(X4))|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X4)|~v1_relat_1(X3)|~v1_relat_1(X4)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_43, c_0_16])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(k5_relat_1(esk3_0,esk2_0)))|~v1_funct_1(k5_relat_1(esk3_0,esk2_0))|~v1_relat_1(k5_relat_1(esk3_0,esk2_0))), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_47, negated_conjecture, (~r2_hidden(esk1_0,k1_relat_1(k5_relat_1(esk3_0,esk2_0)))|~r2_hidden(esk1_0,k1_relat_1(esk3_0))|~r2_hidden(k1_funct_1(esk3_0,esk1_0),k1_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_48, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk1_0),k1_relat_1(X2))|k5_relat_1(esk3_0,esk2_0)!=k5_relat_1(X1,X2)|~v1_funct_1(k5_relat_1(esk3_0,esk2_0))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(k5_relat_1(esk3_0,esk2_0))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_49, negated_conjecture, (~r2_hidden(esk1_0,k1_relat_1(k5_relat_1(esk3_0,esk2_0)))|~r2_hidden(k1_funct_1(esk3_0,esk1_0),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_42])])).
cnf(c_0_50, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk1_0),k1_relat_1(esk2_0))|~v1_funct_1(k5_relat_1(esk3_0,esk2_0))|~v1_relat_1(k5_relat_1(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_48]), c_0_33]), c_0_34]), c_0_35]), c_0_36])])).
cnf(c_0_51, negated_conjecture, (~v1_funct_1(k5_relat_1(esk3_0,esk2_0))|~v1_relat_1(k5_relat_1(esk3_0,esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_46])).
cnf(c_0_52, negated_conjecture, (~v1_relat_1(k5_relat_1(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_21]), c_0_33]), c_0_34]), c_0_35]), c_0_36])])).
cnf(c_0_53, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_20]), c_0_35]), c_0_36])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 54
# Proof object clause steps            : 41
# Proof object formula steps           : 13
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 36
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 266
# ...of these trivial                  : 2
# ...subsumed                          : 50
# ...remaining for further processing  : 214
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 32
# Backward-rewritten                   : 14
# Generated clauses                    : 742
# ...of the previous two non-trivial   : 704
# Contextual simplify-reflections      : 31
# Paramodulations                      : 682
# Factorizations                       : 0
# Equation resolutions                 : 60
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 148
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 142
# Current number of unprocessed clauses: 381
# ...number of literals in the above   : 4832
# Current number of archived formulas  : 0
# Current number of archived clauses   : 66
# Clause-clause subsumption calls (NU) : 10400
# Rec. Clause-clause subsumption calls : 1387
# Non-unit clause-clause subsumptions  : 112
# Unit Clause-clause subsumption calls : 33
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 28451

# -------------------------------------------------
# User time                : 0.033 s
# System time              : 0.000 s
# Total time               : 0.033 s
# Maximum resident set size: 28680 pages
