# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t16_funct_1, conjecture, ![X1]:(![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X2)=X1&k1_relat_1(X3)=X1)=>X2=X3)))=>X1=k1_xboole_0), file('data/MPTP/proofs/27e22b3225db41b5b81630523b9bbd8a.E_input', t16_funct_1)).
fof(s3_funct_1__e7_25__funct_1, axiom, ![X1]:?[X2]:((v1_relat_1(X2)&v1_funct_1(X2))&(k1_relat_1(X2)=X1&![X3]:(r2_hidden(X3,X1)=>k1_funct_1(X2,X3)=1))), file('data/MPTP/proofs/27e22b3225db41b5b81630523b9bbd8a.E_input', s3_funct_1__e7_25__funct_1)).
fof(t64_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>((k1_relat_1(X1)=k1_xboole_0|k2_relat_1(X1)=k1_xboole_0)=>X1=k1_xboole_0)), file('data/MPTP/proofs/27e22b3225db41b5b81630523b9bbd8a.E_input', t64_relat_1)).
fof(s3_funct_1__e2_24__funct_1, axiom, ![X1, X2]:?[X3]:((v1_relat_1(X3)&v1_funct_1(X3))&(k1_relat_1(X3)=X1&![X4]:(r2_hidden(X4,X1)=>k1_funct_1(X3,X4)=X2))), file('data/MPTP/proofs/27e22b3225db41b5b81630523b9bbd8a.E_input', s3_funct_1__e2_24__funct_1)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/27e22b3225db41b5b81630523b9bbd8a.E_input', d4_relat_1)).
fof(t60_relat_1, axiom, (k1_relat_1(k1_xboole_0)=k1_xboole_0&k2_relat_1(k1_xboole_0)=k1_xboole_0), file('data/MPTP/proofs/27e22b3225db41b5b81630523b9bbd8a.E_input', t60_relat_1)).
fof(c_0_6, negated_conjecture, ~(![X1]:(![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X2)=X1&k1_relat_1(X3)=X1)=>X2=X3)))=>X1=k1_xboole_0)), inference(assume_negation,[status(cth)],[t16_funct_1])).
fof(c_0_7, negated_conjecture, ![X6, X7]:((~v1_relat_1(X6)|~v1_funct_1(X6)|(~v1_relat_1(X7)|~v1_funct_1(X7)|(k1_relat_1(X6)!=esk1_0|k1_relat_1(X7)!=esk1_0|X6=X7)))&esk1_0!=k1_xboole_0), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
fof(c_0_8, plain, ![X24, X26]:((v1_relat_1(esk6_1(X24))&v1_funct_1(esk6_1(X24)))&(k1_relat_1(esk6_1(X24))=X24&(~r2_hidden(X26,X24)|k1_funct_1(esk6_1(X24),X26)=1))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[s3_funct_1__e7_25__funct_1])])])])).
fof(c_0_9, plain, ![X23]:((k1_relat_1(X23)!=k1_xboole_0|X23=k1_xboole_0|~v1_relat_1(X23))&(k2_relat_1(X23)!=k1_xboole_0|X23=k1_xboole_0|~v1_relat_1(X23))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t64_relat_1])])])).
fof(c_0_10, plain, ![X8, X9, X11]:((v1_relat_1(esk2_2(X8,X9))&v1_funct_1(esk2_2(X8,X9)))&(k1_relat_1(esk2_2(X8,X9))=X8&(~r2_hidden(X11,X8)|k1_funct_1(esk2_2(X8,X9),X11)=X9))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[s3_funct_1__e2_24__funct_1])])])])).
cnf(c_0_11, negated_conjecture, (X1=X2|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|k1_relat_1(X1)!=esk1_0|k1_relat_1(X2)!=esk1_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (v1_funct_1(esk6_1(X1))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (k1_relat_1(esk6_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, plain, (v1_relat_1(esk6_1(X1))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15, plain, (X1=k1_xboole_0|k1_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, plain, (v1_relat_1(esk2_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, plain, (k1_relat_1(esk2_2(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, negated_conjecture, (X1=esk6_1(esk1_0)|k1_relat_1(X1)!=esk1_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13]), c_0_14])])])).
cnf(c_0_19, plain, (v1_funct_1(esk2_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_20, plain, (k1_funct_1(esk2_2(X2,X3),X1)=X3|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_21, plain, (esk2_2(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])])).
cnf(c_0_22, negated_conjecture, (esk2_2(esk1_0,X1)=esk6_1(esk1_0)), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_17]), c_0_16])])])).
cnf(c_0_23, plain, (k1_funct_1(k1_xboole_0,X1)=X2|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
fof(c_0_24, plain, ![X12, X13, X14, X16, X17, X18, X19, X21]:(((~r2_hidden(X14,X13)|r2_hidden(k4_tarski(X14,esk3_3(X12,X13,X14)),X12)|X13!=k1_relat_1(X12))&(~r2_hidden(k4_tarski(X16,X17),X12)|r2_hidden(X16,X13)|X13!=k1_relat_1(X12)))&((~r2_hidden(esk4_2(X18,X19),X19)|~r2_hidden(k4_tarski(esk4_2(X18,X19),X21),X18)|X19=k1_relat_1(X18))&(r2_hidden(esk4_2(X18,X19),X19)|r2_hidden(k4_tarski(esk4_2(X18,X19),esk5_2(X18,X19)),X18)|X19=k1_relat_1(X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
cnf(c_0_25, negated_conjecture, (k1_funct_1(esk6_1(esk1_0),X1)=X2|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_20, c_0_22])).
cnf(c_0_26, plain, (X1=X2|~r2_hidden(X3,k1_xboole_0)), inference(spm,[status(thm)],[c_0_23, c_0_23])).
cnf(c_0_27, plain, (r2_hidden(esk4_2(X1,X2),X2)|r2_hidden(k4_tarski(esk4_2(X1,X2),esk5_2(X1,X2)),X1)|X2=k1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_28, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
cnf(c_0_29, negated_conjecture, (X1=X2|~r2_hidden(X3,esk1_0)), inference(spm,[status(thm)],[c_0_25, c_0_25])).
cnf(c_0_30, plain, (X1=k1_xboole_0|X2=X3|r2_hidden(esk4_2(k1_xboole_0,X1),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])).
cnf(c_0_31, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_32, negated_conjecture, (X1=X2|X3=X4), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])).
cnf(c_0_33, negated_conjecture, (X1=X2), inference(er,[status(thm)],[inference(ef,[status(thm)],[c_0_32])])).
cnf(c_0_34, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_31, c_0_33]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 35
# Proof object clause steps            : 23
# Proof object formula steps           : 12
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 14
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 18
# Processed clauses                    : 75
# ...of these trivial                  : 0
# ...subsumed                          : 6
# ...remaining for further processing  : 69
# Other redundant clauses eliminated   : 32
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 35
# Backward-rewritten                   : 0
# Generated clauses                    : 886
# ...of the previous two non-trivial   : 813
# Contextual simplify-reflections      : 0
# Paramodulations                      : 841
# Factorizations                       : 12
# Equation resolutions                 : 32
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 11
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 4
# Current number of unprocessed clauses: 742
# ...number of literals in the above   : 2366
# Current number of archived formulas  : 0
# Current number of archived clauses   : 54
# Clause-clause subsumption calls (NU) : 94
# Rec. Clause-clause subsumption calls : 91
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 40
# Rewrite failures with RHS unbound    : 92
# BW rewrite match attempts            : 84
# BW rewrite match successes           : 84
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5964

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.003 s
# Total time               : 0.014 s
# Maximum resident set size: 28596 pages
