# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C41_F1_PI_AE_Q4_CS_SP_PS_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t44_subset_1, conjecture, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(r1_xboole_0(X2,k3_subset_1(X1,X3))<=>r1_tarski(X2,X3)))), file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', t44_subset_1)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', d1_zfmisc_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', d2_subset_1)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', fc1_subset_1)).
fof(t79_xboole_1, axiom, ![X1, X2]:r1_xboole_0(k4_xboole_0(X1,X2),X2), file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', t79_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', t100_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', commutativity_k3_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', t28_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', t3_boole)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', d7_xboole_0)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', t65_xboole_1)).
fof(d5_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)), file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', d5_subset_1)).
fof(t16_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3)), file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', t16_xboole_1)).
fof(t106_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k4_xboole_0(X2,X3))=>(r1_tarski(X1,X2)&r1_xboole_0(X1,X3))), file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', t106_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', t48_xboole_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', idempotence_k3_xboole_0)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', t41_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', d10_xboole_0)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', commutativity_k2_xboole_0)).
fof(t73_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,k2_xboole_0(X2,X3))&r1_xboole_0(X1,X3))=>r1_tarski(X1,X2)), file('data/MPTP/proofs/27e263b5da094864b3912b0338761c9e.E_input', t73_xboole_1)).
fof(c_0_20, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(r1_xboole_0(X2,k3_subset_1(X1,X3))<=>r1_tarski(X2,X3))))), inference(assume_negation,[status(cth)],[t44_subset_1])).
fof(c_0_21, plain, ![X25, X26, X27, X28, X29, X30]:(((~r2_hidden(X27,X26)|r1_tarski(X27,X25)|X26!=k1_zfmisc_1(X25))&(~r1_tarski(X28,X25)|r2_hidden(X28,X26)|X26!=k1_zfmisc_1(X25)))&((~r2_hidden(esk4_2(X29,X30),X30)|~r1_tarski(esk4_2(X29,X30),X29)|X30=k1_zfmisc_1(X29))&(r2_hidden(esk4_2(X29,X30),X30)|r1_tarski(esk4_2(X29,X30),X29)|X30=k1_zfmisc_1(X29)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
fof(c_0_22, plain, ![X48, X49]:(((~m1_subset_1(X49,X48)|r2_hidden(X49,X48)|v1_xboole_0(X48))&(~r2_hidden(X49,X48)|m1_subset_1(X49,X48)|v1_xboole_0(X48)))&((~m1_subset_1(X49,X48)|v1_xboole_0(X49)|~v1_xboole_0(X48))&(~v1_xboole_0(X49)|m1_subset_1(X49,X48)|~v1_xboole_0(X48)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_23, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))&((~r1_xboole_0(esk2_0,k3_subset_1(esk1_0,esk3_0))|~r1_tarski(esk2_0,esk3_0))&(r1_xboole_0(esk2_0,k3_subset_1(esk1_0,esk3_0))|r1_tarski(esk2_0,esk3_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).
fof(c_0_24, plain, ![X12]:~v1_xboole_0(k1_zfmisc_1(X12)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
fof(c_0_25, plain, ![X8, X9]:r1_xboole_0(k4_xboole_0(X8,X9),X9), inference(variable_rename,[status(thm)],[t79_xboole_1])).
fof(c_0_26, plain, ![X43, X44]:k4_xboole_0(X43,X44)=k5_xboole_0(X43,k3_xboole_0(X43,X44)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_27, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, plain, (r1_xboole_0(k4_xboole_0(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_33, plain, ![X15, X16]:k3_xboole_0(X15,X16)=k3_xboole_0(X16,X15), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_34, plain, ![X34, X35]:(~r1_tarski(X34,X35)|k3_xboole_0(X34,X35)=X34), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_35, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_36, negated_conjecture, (r2_hidden(esk3_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])).
cnf(c_0_37, plain, (r1_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X2)), inference(rw,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_38, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_39, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_40, negated_conjecture, (r1_tarski(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
fof(c_0_41, plain, ![X42]:k4_xboole_0(X42,k1_xboole_0)=X42, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_42, plain, ![X23, X24]:((~r1_xboole_0(X23,X24)|k3_xboole_0(X23,X24)=k1_xboole_0)&(k3_xboole_0(X23,X24)!=k1_xboole_0|r1_xboole_0(X23,X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
fof(c_0_43, plain, ![X22]:r1_xboole_0(X22,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
fof(c_0_44, plain, ![X32, X33]:(~m1_subset_1(X33,k1_zfmisc_1(X32))|k3_subset_1(X32,X33)=k4_xboole_0(X32,X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
fof(c_0_45, plain, ![X45, X46, X47]:k3_xboole_0(k3_xboole_0(X45,X46),X47)=k3_xboole_0(X45,k3_xboole_0(X46,X47)), inference(variable_rename,[status(thm)],[t16_xboole_1])).
cnf(c_0_46, plain, (r1_xboole_0(k5_xboole_0(X1,k3_xboole_0(X2,X1)),X2)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_47, negated_conjecture, (k3_xboole_0(esk3_0,esk1_0)=esk3_0), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_48, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_49, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_50, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_51, plain, ![X36, X37, X38]:((r1_tarski(X36,X37)|~r1_tarski(X36,k4_xboole_0(X37,X38)))&(r1_xboole_0(X36,X38)|~r1_tarski(X36,k4_xboole_0(X37,X38)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t106_xboole_1])])])).
fof(c_0_52, plain, ![X10, X11]:k4_xboole_0(X10,k4_xboole_0(X10,X11))=k3_xboole_0(X10,X11), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_53, plain, ![X7]:k3_xboole_0(X7,X7)=X7, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_54, plain, (k3_subset_1(X2,X1)=k4_xboole_0(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_55, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_56, negated_conjecture, (r1_xboole_0(k5_xboole_0(esk1_0,esk3_0),esk3_0)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_57, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[c_0_48, c_0_32])).
cnf(c_0_58, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_59, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k4_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_60, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_61, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_62, negated_conjecture, (r1_xboole_0(esk2_0,k3_subset_1(esk1_0,esk3_0))|r1_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_63, plain, (k5_xboole_0(X2,k3_xboole_0(X2,X1))=k3_subset_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_54, c_0_32])).
cnf(c_0_64, plain, (r1_xboole_0(k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,k3_xboole_0(X2,X3))),X3)), inference(spm,[status(thm)],[c_0_37, c_0_55])).
cnf(c_0_65, negated_conjecture, (k3_xboole_0(esk3_0,k5_xboole_0(esk1_0,esk3_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_56]), c_0_38])).
cnf(c_0_66, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_57, c_0_58])).
fof(c_0_67, plain, ![X17, X18, X19]:k4_xboole_0(k4_xboole_0(X17,X18),X19)=k4_xboole_0(X17,k2_xboole_0(X18,X19)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
cnf(c_0_68, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))), inference(rw,[status(thm)],[c_0_59, c_0_32])).
fof(c_0_69, plain, ![X13, X14]:(((r1_tarski(X13,X14)|X13!=X14)&(r1_tarski(X14,X13)|X13!=X14))&(~r1_tarski(X13,X14)|~r1_tarski(X14,X13)|X13=X14)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_70, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2))))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_32]), c_0_32])).
cnf(c_0_71, plain, (k3_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_55, c_0_61])).
cnf(c_0_72, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k4_xboole_0(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_73, negated_conjecture, (k3_xboole_0(esk2_0,k3_subset_1(esk1_0,esk3_0))=k1_xboole_0|r1_tarski(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_49, c_0_62])).
cnf(c_0_74, negated_conjecture, (k3_subset_1(esk1_0,esk3_0)=k5_xboole_0(esk1_0,k3_xboole_0(esk3_0,esk1_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_29]), c_0_38])).
cnf(c_0_75, negated_conjecture, (r1_xboole_0(k3_xboole_0(X1,esk3_0),k5_xboole_0(esk1_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_58]), c_0_66])).
cnf(c_0_76, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_77, plain, (k3_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1)))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_37]), c_0_38])).
cnf(c_0_78, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k5_xboole_0(X2,k3_xboole_0(X3,X2)))), inference(spm,[status(thm)],[c_0_68, c_0_38])).
cnf(c_0_79, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_80, plain, (k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2)))=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_70]), c_0_71])).
cnf(c_0_81, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k5_xboole_0(X3,k3_xboole_0(X3,X2)))), inference(rw,[status(thm)],[c_0_72, c_0_32])).
cnf(c_0_82, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(X3,k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_38, c_0_55])).
cnf(c_0_83, negated_conjecture, (k3_xboole_0(esk2_0,k5_xboole_0(esk1_0,esk3_0))=k1_xboole_0|r1_tarski(esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_74]), c_0_47])).
cnf(c_0_84, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_38, c_0_58])).
cnf(c_0_85, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_86, negated_conjecture, (r1_xboole_0(k3_xboole_0(esk3_0,X1),k5_xboole_0(esk1_0,esk3_0))), inference(spm,[status(thm)],[c_0_75, c_0_38])).
cnf(c_0_87, plain, (k5_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),k3_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X3))=k5_xboole_0(X1,k3_xboole_0(X1,k2_xboole_0(X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_32]), c_0_32]), c_0_32])).
cnf(c_0_88, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_77]), c_0_66]), c_0_61])).
cnf(c_0_89, negated_conjecture, (r1_tarski(X1,esk1_0)|~r1_tarski(X1,k5_xboole_0(esk1_0,esk3_0))), inference(spm,[status(thm)],[c_0_78, c_0_47])).
cnf(c_0_90, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_79])).
cnf(c_0_91, plain, (k5_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2)))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_70, c_0_80])).
cnf(c_0_92, plain, (r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_tarski(X1,k5_xboole_0(X4,k3_xboole_0(X2,k3_xboole_0(X3,X4))))), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_93, negated_conjecture, (k3_xboole_0(esk2_0,k3_xboole_0(k5_xboole_0(esk1_0,esk3_0),X1))=k1_xboole_0|r1_tarski(esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_83]), c_0_84])).
cnf(c_0_94, negated_conjecture, (r2_hidden(esk2_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_85]), c_0_30])).
cnf(c_0_95, negated_conjecture, (k3_xboole_0(esk3_0,k3_xboole_0(X1,k5_xboole_0(esk1_0,esk3_0)))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_86]), c_0_55])).
cnf(c_0_96, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k2_xboole_0(X2,k5_xboole_0(X1,k3_xboole_0(X1,X2)))))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_61]), c_0_88])).
cnf(c_0_97, negated_conjecture, (r1_tarski(k5_xboole_0(esk1_0,esk3_0),esk1_0)), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_98, plain, (k5_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X2,X1)))=k3_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_91, c_0_38])).
fof(c_0_99, plain, ![X20, X21]:k2_xboole_0(X20,X21)=k2_xboole_0(X21,X20), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_100, negated_conjecture, (r1_tarski(esk2_0,esk3_0)|r1_xboole_0(X1,k3_xboole_0(esk2_0,k5_xboole_0(esk1_0,esk3_0)))|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_66])).
cnf(c_0_101, negated_conjecture, (r1_tarski(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_35, c_0_94])).
cnf(c_0_102, negated_conjecture, (r1_xboole_0(X1,k3_xboole_0(X2,k5_xboole_0(esk1_0,esk3_0)))|~r1_tarski(X1,esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_95]), c_0_66])).
cnf(c_0_103, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k3_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_68, c_0_70])).
cnf(c_0_104, plain, (k3_xboole_0(X1,k2_xboole_0(X2,k5_xboole_0(X1,k3_xboole_0(X1,X2))))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_96]), c_0_66])).
cnf(c_0_105, negated_conjecture, (k3_xboole_0(esk1_0,k5_xboole_0(esk1_0,esk3_0))=k5_xboole_0(esk1_0,esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_97]), c_0_38])).
cnf(c_0_106, negated_conjecture, (k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,esk3_0))=esk3_0), inference(spm,[status(thm)],[c_0_98, c_0_47])).
cnf(c_0_107, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_108, negated_conjecture, (r1_xboole_0(esk2_0,k3_xboole_0(esk2_0,k5_xboole_0(esk1_0,esk3_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_102])).
cnf(c_0_109, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k3_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_103, c_0_38])).
cnf(c_0_110, negated_conjecture, (k3_xboole_0(esk1_0,k2_xboole_0(esk3_0,k5_xboole_0(esk1_0,esk3_0)))=esk1_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_106]), c_0_107])).
cnf(c_0_111, negated_conjecture, (~r1_xboole_0(esk2_0,k3_subset_1(esk1_0,esk3_0))|~r1_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_112, negated_conjecture, (k3_xboole_0(esk2_0,k5_xboole_0(esk1_0,esk3_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_108]), c_0_71])).
fof(c_0_113, plain, ![X39, X40, X41]:(~r1_tarski(X39,k2_xboole_0(X40,X41))|~r1_xboole_0(X39,X41)|r1_tarski(X39,X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t73_xboole_1])])).
cnf(c_0_114, negated_conjecture, (r1_tarski(X1,k2_xboole_0(esk3_0,k5_xboole_0(esk1_0,esk3_0)))|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_109, c_0_110])).
cnf(c_0_115, negated_conjecture, (~r1_tarski(esk2_0,esk3_0)|~r1_xboole_0(esk2_0,k5_xboole_0(esk1_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_111, c_0_74]), c_0_47])).
cnf(c_0_116, negated_conjecture, (r1_xboole_0(esk2_0,k5_xboole_0(esk1_0,esk3_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_112]), c_0_66])).
cnf(c_0_117, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_xboole_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_118, negated_conjecture, (r1_tarski(esk2_0,k2_xboole_0(esk3_0,k5_xboole_0(esk1_0,esk3_0)))), inference(spm,[status(thm)],[c_0_114, c_0_101])).
cnf(c_0_119, negated_conjecture, (~r1_tarski(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_115, c_0_116])])).
cnf(c_0_120, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_116])]), c_0_119]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 121
# Proof object clause steps            : 80
# Proof object formula steps           : 41
# Proof object conjectures             : 36
# Proof object clause conjectures      : 33
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 43
# Proof object simplifying inferences  : 44
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 20
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 32
# Processed clauses                    : 7306
# ...of these trivial                  : 1266
# ...subsumed                          : 4566
# ...remaining for further processing  : 1474
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 26
# Backward-rewritten                   : 160
# Generated clauses                    : 93325
# ...of the previous two non-trivial   : 63116
# Contextual simplify-reflections      : 1
# Paramodulations                      : 93306
# Factorizations                       : 4
# Equation resolutions                 : 9
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1249
#    Positive orientable unit clauses  : 637
#    Positive unorientable unit clauses: 6
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 604
# Current number of unprocessed clauses: 55389
# ...number of literals in the above   : 76854
# Current number of archived formulas  : 0
# Current number of archived clauses   : 224
# Clause-clause subsumption calls (NU) : 94600
# Rec. Clause-clause subsumption calls : 76087
# Non-unit clause-clause subsumptions  : 4457
# Unit Clause-clause subsumption calls : 1704
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2982
# BW rewrite match successes           : 175
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1308187

# -------------------------------------------------
# User time                : 0.615 s
# System time              : 0.027 s
# Total time               : 0.642 s
# Maximum resident set size: 54596 pages
