# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/284072baa5544650ac78def596a879c8.E_input', d3_tarski)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/284072baa5544650ac78def596a879c8.E_input', t2_xboole_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/284072baa5544650ac78def596a879c8.E_input', d5_xboole_0)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/284072baa5544650ac78def596a879c8.E_input', d1_zfmisc_1)).
fof(t1_zfmisc_1, conjecture, k1_zfmisc_1(k1_xboole_0)=k1_tarski(k1_xboole_0), file('data/MPTP/proofs/284072baa5544650ac78def596a879c8.E_input', t1_zfmisc_1)).
fof(t3_xboole_1, axiom, ![X1]:(r1_tarski(X1,k1_xboole_0)=>X1=k1_xboole_0), file('data/MPTP/proofs/284072baa5544650ac78def596a879c8.E_input', t3_xboole_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/284072baa5544650ac78def596a879c8.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/284072baa5544650ac78def596a879c8.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/284072baa5544650ac78def596a879c8.E_input', t71_enumset1)).
fof(d1_enumset1, axiom, ![X1, X2, X3, X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3)))))), file('data/MPTP/proofs/284072baa5544650ac78def596a879c8.E_input', d1_enumset1)).
fof(c_0_10, plain, ![X12, X13, X14, X15, X16]:((~r1_tarski(X12,X13)|(~r2_hidden(X14,X12)|r2_hidden(X14,X13)))&((r2_hidden(esk1_2(X15,X16),X15)|r1_tarski(X15,X16))&(~r2_hidden(esk1_2(X15,X16),X16)|r1_tarski(X15,X16)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_11, plain, ![X9]:r1_tarski(k1_xboole_0,X9), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_12, plain, ![X38, X39, X40, X41, X42, X43, X44, X45]:((((r2_hidden(X41,X38)|~r2_hidden(X41,X40)|X40!=k4_xboole_0(X38,X39))&(~r2_hidden(X41,X39)|~r2_hidden(X41,X40)|X40!=k4_xboole_0(X38,X39)))&(~r2_hidden(X42,X38)|r2_hidden(X42,X39)|r2_hidden(X42,X40)|X40!=k4_xboole_0(X38,X39)))&((~r2_hidden(esk4_3(X43,X44,X45),X45)|(~r2_hidden(esk4_3(X43,X44,X45),X43)|r2_hidden(esk4_3(X43,X44,X45),X44))|X45=k4_xboole_0(X43,X44))&((r2_hidden(esk4_3(X43,X44,X45),X43)|r2_hidden(esk4_3(X43,X44,X45),X45)|X45=k4_xboole_0(X43,X44))&(~r2_hidden(esk4_3(X43,X44,X45),X44)|r2_hidden(esk4_3(X43,X44,X45),X45)|X45=k4_xboole_0(X43,X44))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_13, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (r2_hidden(esk4_3(X1,X2,X3),X1)|r2_hidden(esk4_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_18, plain, (k4_xboole_0(X1,X2)=X1|r2_hidden(esk4_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_15])).
cnf(c_0_19, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0|r2_hidden(esk4_3(k1_xboole_0,X1,k1_xboole_0),X2)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_21, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_20])).
cnf(c_0_22, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23, plain, (~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_21]), c_0_17])).
fof(c_0_24, plain, ![X29, X30, X31, X32, X33, X34]:(((~r2_hidden(X31,X30)|r1_tarski(X31,X29)|X30!=k1_zfmisc_1(X29))&(~r1_tarski(X32,X29)|r2_hidden(X32,X30)|X30!=k1_zfmisc_1(X29)))&((~r2_hidden(esk3_2(X33,X34),X34)|~r1_tarski(esk3_2(X33,X34),X33)|X34=k1_zfmisc_1(X33))&(r2_hidden(esk3_2(X33,X34),X34)|r1_tarski(esk3_2(X33,X34),X33)|X34=k1_zfmisc_1(X33)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
cnf(c_0_25, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_21]), c_0_23])).
cnf(c_0_26, plain, (r2_hidden(esk1_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_27, plain, (r2_hidden(X1,X3)|~r1_tarski(X1,X2)|X3!=k1_zfmisc_1(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_28, plain, (r1_tarski(X1,X2)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_29, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_30, negated_conjecture, ~(k1_zfmisc_1(k1_xboole_0)=k1_tarski(k1_xboole_0)), inference(assume_negation,[status(cth)],[t1_zfmisc_1])).
cnf(c_0_31, plain, (r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_27, c_0_28])).
fof(c_0_32, plain, ![X36]:(~r1_tarski(X36,k1_xboole_0)|X36=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_xboole_1])])).
cnf(c_0_33, plain, (r1_tarski(esk1_2(X1,X2),X3)|r1_tarski(X1,X2)|X1!=k1_zfmisc_1(X3)), inference(spm,[status(thm)],[c_0_29, c_0_26])).
fof(c_0_34, negated_conjecture, k1_zfmisc_1(k1_xboole_0)!=k1_tarski(k1_xboole_0), inference(fof_simplification,[status(thm)],[c_0_30])).
fof(c_0_35, plain, ![X37]:k2_tarski(X37,X37)=k1_tarski(X37), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_36, plain, ![X10, X11]:k1_enumset1(X10,X10,X11)=k2_tarski(X10,X11), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_37, plain, ![X6, X7, X8]:k2_enumset1(X6,X6,X7,X8)=k1_enumset1(X6,X7,X8), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_38, plain, ![X18, X19, X20, X21, X22, X23, X24, X25, X26, X27]:(((~r2_hidden(X22,X21)|(X22=X18|(X22=X19|X22=X20))|X21!=k1_enumset1(X18,X19,X20))&((X23!=X18|r2_hidden(X23,X21)|X21!=k1_enumset1(X18,X19,X20))&((X23!=X19|r2_hidden(X23,X21)|X21!=k1_enumset1(X18,X19,X20))&(X23!=X20|r2_hidden(X23,X21)|X21!=k1_enumset1(X18,X19,X20)))))&(((esk2_4(X24,X25,X26,X27)!=X24|~r2_hidden(esk2_4(X24,X25,X26,X27),X27)|X27=k1_enumset1(X24,X25,X26))&((esk2_4(X24,X25,X26,X27)!=X25|~r2_hidden(esk2_4(X24,X25,X26,X27),X27)|X27=k1_enumset1(X24,X25,X26))&(esk2_4(X24,X25,X26,X27)!=X26|~r2_hidden(esk2_4(X24,X25,X26,X27),X27)|X27=k1_enumset1(X24,X25,X26))))&(r2_hidden(esk2_4(X24,X25,X26,X27),X27)|(esk2_4(X24,X25,X26,X27)=X24|(esk2_4(X24,X25,X26,X27)=X25|esk2_4(X24,X25,X26,X27)=X26))|X27=k1_enumset1(X24,X25,X26)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])).
cnf(c_0_39, plain, (r1_tarski(X1,X2)|~r2_hidden(esk1_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_40, plain, (r2_hidden(X1,k1_zfmisc_1(X2))|X1!=k1_xboole_0), inference(er,[status(thm)],[c_0_31])).
cnf(c_0_41, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_42, plain, (r1_tarski(esk1_2(k1_zfmisc_1(X1),X2),X1)|r1_tarski(k1_zfmisc_1(X1),X2)), inference(er,[status(thm)],[c_0_33])).
cnf(c_0_43, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)!=k1_tarski(k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_44, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_45, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_46, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_47, plain, (r2_hidden(esk2_4(X1,X2,X3,X4),X4)|esk2_4(X1,X2,X3,X4)=X1|esk2_4(X1,X2,X3,X4)=X2|esk2_4(X1,X2,X3,X4)=X3|X4=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_48, plain, (r1_tarski(X1,k1_zfmisc_1(X2))|esk1_2(X1,k1_zfmisc_1(X2))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_49, plain, (esk1_2(k1_zfmisc_1(k1_xboole_0),X1)=k1_xboole_0|r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_50, negated_conjecture, (k2_enumset1(k1_xboole_0,k1_xboole_0,k1_xboole_0,k1_xboole_0)!=k1_zfmisc_1(k1_xboole_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_46])).
cnf(c_0_51, plain, (X4=k2_enumset1(X1,X1,X2,X3)|esk2_4(X1,X2,X3,X4)=X3|esk2_4(X1,X2,X3,X4)=X2|esk2_4(X1,X2,X3,X4)=X1|r2_hidden(esk2_4(X1,X2,X3,X4),X4)), inference(rw,[status(thm)],[c_0_47, c_0_46])).
cnf(c_0_52, plain, (r1_tarski(k1_zfmisc_1(k1_xboole_0),k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_53, negated_conjecture, (esk2_4(k1_xboole_0,k1_xboole_0,k1_xboole_0,X1)=k1_xboole_0|r2_hidden(esk2_4(k1_xboole_0,k1_xboole_0,k1_xboole_0,X1),X1)|X1!=k1_zfmisc_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_54, plain, (r2_hidden(X1,k1_zfmisc_1(X2))|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_13, c_0_52])).
cnf(c_0_55, negated_conjecture, (esk2_4(k1_xboole_0,k1_xboole_0,k1_xboole_0,k1_zfmisc_1(k1_xboole_0))=k1_xboole_0|r2_hidden(esk2_4(k1_xboole_0,k1_xboole_0,k1_xboole_0,k1_zfmisc_1(k1_xboole_0)),k1_zfmisc_1(k1_xboole_0))), inference(er,[status(thm)],[c_0_53])).
cnf(c_0_56, negated_conjecture, (r2_hidden(esk2_4(k1_xboole_0,k1_xboole_0,k1_xboole_0,k1_zfmisc_1(k1_xboole_0)),k1_zfmisc_1(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_40])).
cnf(c_0_57, plain, (X4=k1_enumset1(X1,X2,X3)|esk2_4(X1,X2,X3,X4)!=X3|~r2_hidden(esk2_4(X1,X2,X3,X4),X4)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_58, negated_conjecture, (r1_tarski(esk2_4(k1_xboole_0,k1_xboole_0,k1_xboole_0,k1_zfmisc_1(k1_xboole_0)),X1)|k1_zfmisc_1(X2)!=k1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_29, c_0_56])).
cnf(c_0_59, plain, (X4=k2_enumset1(X1,X1,X2,X3)|esk2_4(X1,X2,X3,X4)!=X3|~r2_hidden(esk2_4(X1,X2,X3,X4),X4)), inference(rw,[status(thm)],[c_0_57, c_0_46])).
cnf(c_0_60, negated_conjecture, (r1_tarski(esk2_4(k1_xboole_0,k1_xboole_0,k1_xboole_0,k1_zfmisc_1(k1_xboole_0)),X1)), inference(er,[status(thm)],[c_0_58])).
cnf(c_0_61, negated_conjecture, (esk2_4(k1_xboole_0,k1_xboole_0,k1_xboole_0,k1_zfmisc_1(k1_xboole_0))!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_56]), c_0_50])).
cnf(c_0_62, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_60]), c_0_61]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 63
# Proof object clause steps            : 42
# Proof object formula steps           : 21
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 23
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 26
# Processed clauses                    : 463
# ...of these trivial                  : 8
# ...subsumed                          : 208
# ...remaining for further processing  : 247
# Other redundant clauses eliminated   : 403
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 13
# Backward-rewritten                   : 15
# Generated clauses                    : 4730
# ...of the previous two non-trivial   : 4075
# Contextual simplify-reflections      : 7
# Paramodulations                      : 4245
# Factorizations                       : 27
# Equation resolutions                 : 455
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 189
#    Positive orientable unit clauses  : 24
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 8
#    Non-unit-clauses                  : 157
# Current number of unprocessed clauses: 3616
# ...number of literals in the above   : 28720
# Current number of archived formulas  : 0
# Current number of archived clauses   : 57
# Clause-clause subsumption calls (NU) : 4262
# Rec. Clause-clause subsumption calls : 2275
# Non-unit clause-clause subsumptions  : 105
# Unit Clause-clause subsumption calls : 381
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 36
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 89207

# -------------------------------------------------
# User time                : 0.112 s
# System time              : 0.000 s
# Total time               : 0.112 s
# Maximum resident set size: 28716 pages
