# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t39_tops_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k2_pre_topc(X1,X2))<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>~((v3_pre_topc(X4,X1)&(r2_hidden(X3,X4)&r1_xboole_0(X2,X4))))))))), file('data/MPTP/proofs/28c05c59a0b049f5ad6a55e9f8514bd8.E_input', t39_tops_1)).
fof(t54_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k2_pre_topc(X1,X2))<=>(~(v2_struct_0(X1))&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>~((v3_pre_topc(X4,X1)&(r2_hidden(X3,X4)&r1_xboole_0(X2,X4)))))))))), file('data/MPTP/proofs/28c05c59a0b049f5ad6a55e9f8514bd8.E_input', t54_pre_topc)).
fof(c_0_2, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k2_pre_topc(X1,X2))<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>~((v3_pre_topc(X4,X1)&(r2_hidden(X3,X4)&r1_xboole_0(X2,X4)))))))))), inference(assume_negation,[status(cth)],[t39_tops_1])).
fof(c_0_3, plain, ![X10, X11, X12, X13]:(((~v2_struct_0(X10)|~r2_hidden(X12,k2_pre_topc(X10,X11))|~m1_subset_1(X12,u1_struct_0(X10))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~l1_pre_topc(X10))&(~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X10)))|(~v3_pre_topc(X13,X10)|(~r2_hidden(X12,X13)|~r1_xboole_0(X11,X13)))|~r2_hidden(X12,k2_pre_topc(X10,X11))|~m1_subset_1(X12,u1_struct_0(X10))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~l1_pre_topc(X10)))&((m1_subset_1(esk5_3(X10,X11,X12),k1_zfmisc_1(u1_struct_0(X10)))|v2_struct_0(X10)|r2_hidden(X12,k2_pre_topc(X10,X11))|~m1_subset_1(X12,u1_struct_0(X10))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~l1_pre_topc(X10))&((v3_pre_topc(esk5_3(X10,X11,X12),X10)|v2_struct_0(X10)|r2_hidden(X12,k2_pre_topc(X10,X11))|~m1_subset_1(X12,u1_struct_0(X10))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~l1_pre_topc(X10))&((r2_hidden(X12,esk5_3(X10,X11,X12))|v2_struct_0(X10)|r2_hidden(X12,k2_pre_topc(X10,X11))|~m1_subset_1(X12,u1_struct_0(X10))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~l1_pre_topc(X10))&(r1_xboole_0(X11,esk5_3(X10,X11,X12))|v2_struct_0(X10)|r2_hidden(X12,k2_pre_topc(X10,X11))|~m1_subset_1(X12,u1_struct_0(X10))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~l1_pre_topc(X10)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t54_pre_topc])])])])])])).
fof(c_0_4, negated_conjecture, ![X9]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(((m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0)))&((v3_pre_topc(esk4_0,esk1_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0)))&((r2_hidden(esk3_0,esk4_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0)))&(r1_xboole_0(esk2_0,esk4_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))))))&(r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|(~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(esk1_0)))|(~v3_pre_topc(X9,esk1_0)|(~r2_hidden(esk3_0,X9)|~r1_xboole_0(esk2_0,X9))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_2])])])])])])).
cnf(c_0_5, plain, (r1_xboole_0(X1,esk5_3(X2,X1,X3))|v2_struct_0(X2)|r2_hidden(X3,k2_pre_topc(X2,X1))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_6, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9, plain, (r2_hidden(X1,esk5_3(X2,X3,X1))|v2_struct_0(X2)|r2_hidden(X1,k2_pre_topc(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_10, plain, (v3_pre_topc(esk5_3(X1,X2,X3),X1)|v2_struct_0(X1)|r2_hidden(X3,k2_pre_topc(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_11, negated_conjecture, (r1_xboole_0(X1,esk5_3(esk1_0,X1,esk3_0))|r2_hidden(esk3_0,k2_pre_topc(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5, c_0_6]), c_0_7])]), c_0_8])).
cnf(c_0_12, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_13, negated_conjecture, (r2_hidden(esk3_0,esk5_3(esk1_0,X1,esk3_0))|r2_hidden(esk3_0,k2_pre_topc(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_6]), c_0_7])]), c_0_8])).
cnf(c_0_14, negated_conjecture, (v3_pre_topc(esk5_3(esk1_0,X1,esk3_0),esk1_0)|r2_hidden(esk3_0,k2_pre_topc(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_6]), c_0_7])]), c_0_8])).
cnf(c_0_15, plain, (m1_subset_1(esk5_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|r2_hidden(X3,k2_pre_topc(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_16, negated_conjecture, (r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v3_pre_topc(X1,esk1_0)|~r2_hidden(esk3_0,X1)|~r1_xboole_0(esk2_0,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_17, negated_conjecture, (r1_xboole_0(esk2_0,esk5_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_18, negated_conjecture, (r2_hidden(esk3_0,esk5_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_13, c_0_12])).
cnf(c_0_19, negated_conjecture, (v3_pre_topc(esk5_3(esk1_0,esk2_0,esk3_0),esk1_0)|r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_14, c_0_12])).
cnf(c_0_20, negated_conjecture, (r2_hidden(esk3_0,k2_pre_topc(esk1_0,X1))|m1_subset_1(esk5_3(esk1_0,X1,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_6]), c_0_7])]), c_0_8])).
cnf(c_0_21, negated_conjecture, (r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|~m1_subset_1(esk5_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18]), c_0_19])).
cnf(c_0_22, plain, (~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v3_pre_topc(X1,X2)|~r2_hidden(X3,X1)|~r1_xboole_0(X4,X1)|~r2_hidden(X3,k2_pre_topc(X2,X4))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_23, negated_conjecture, (v3_pre_topc(esk4_0,esk1_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_24, negated_conjecture, (r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_12]), c_0_21])).
cnf(c_0_25, negated_conjecture, (r2_hidden(esk3_0,esk4_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_26, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_27, negated_conjecture, (~r1_xboole_0(X1,X2)|~v3_pre_topc(X2,esk1_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,X1))|~r2_hidden(esk3_0,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_6]), c_0_7])])).
cnf(c_0_28, negated_conjecture, (v3_pre_topc(esk4_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_24])])).
cnf(c_0_29, negated_conjecture, (r2_hidden(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_24])])).
cnf(c_0_30, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_24])])).
cnf(c_0_31, negated_conjecture, (r1_xboole_0(esk2_0,esk4_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_32, negated_conjecture, (~r1_xboole_0(X1,esk4_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30])])).
cnf(c_0_33, negated_conjecture, (r1_xboole_0(esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_24])])).
cnf(c_0_34, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_24]), c_0_12])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 35
# Proof object clause steps            : 30
# Proof object formula steps           : 5
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 2
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 31
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 2
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 16
# Processed clauses                    : 50
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 49
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 9
# Generated clauses                    : 18
# ...of the previous two non-trivial   : 21
# Contextual simplify-reflections      : 3
# Paramodulations                      : 18
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 24
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 14
# Current number of unprocessed clauses: 2
# ...number of literals in the above   : 4
# Current number of archived formulas  : 0
# Current number of archived clauses   : 25
# Clause-clause subsumption calls (NU) : 86
# Rec. Clause-clause subsumption calls : 47
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 29
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2093

# -------------------------------------------------
# User time                : 0.004 s
# System time              : 0.004 s
# Total time               : 0.009 s
# Maximum resident set size: 28616 pages
