# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t53_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r3_wellord1(X1,X2,X3)=>((v1_relat_2(X1)=>v1_relat_2(X2))&((v8_relat_2(X1)=>v8_relat_2(X2))&((v6_relat_2(X1)=>v6_relat_2(X2))&((v4_relat_2(X1)=>v4_relat_2(X2))&(v1_wellord1(X1)=>v1_wellord1(X2)))))))))), file('data/MPTP/proofs/28d06fed8d674af98c096338e0325c04.E_input', t53_wellord1)).
fof(t54_wellord1, conjecture, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((v2_wellord1(X1)&r3_wellord1(X1,X2,X3))=>v2_wellord1(X2))))), file('data/MPTP/proofs/28d06fed8d674af98c096338e0325c04.E_input', t54_wellord1)).
fof(d4_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v2_wellord1(X1)<=>(v1_relat_2(X1)&(v8_relat_2(X1)&(v4_relat_2(X1)&(v6_relat_2(X1)&v1_wellord1(X1))))))), file('data/MPTP/proofs/28d06fed8d674af98c096338e0325c04.E_input', d4_wellord1)).
fof(c_0_3, plain, ![X1, X2]:(epred1_2(X2,X1)<=>((v1_relat_2(X1)=>v1_relat_2(X2))&((v8_relat_2(X1)=>v8_relat_2(X2))&((v6_relat_2(X1)=>v6_relat_2(X2))&((v4_relat_2(X1)=>v4_relat_2(X2))&(v1_wellord1(X1)=>v1_wellord1(X2))))))), introduced(definition)).
fof(c_0_4, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r3_wellord1(X1,X2,X3)=>epred1_2(X2,X1))))), inference(apply_def,[status(thm)],[t53_wellord1, c_0_3])).
fof(c_0_5, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((v2_wellord1(X1)&r3_wellord1(X1,X2,X3))=>v2_wellord1(X2)))))), inference(assume_negation,[status(cth)],[t54_wellord1])).
fof(c_0_6, plain, ![X1, X2]:(epred1_2(X2,X1)=>((v1_relat_2(X1)=>v1_relat_2(X2))&((v8_relat_2(X1)=>v8_relat_2(X2))&((v6_relat_2(X1)=>v6_relat_2(X2))&((v4_relat_2(X1)=>v4_relat_2(X2))&(v1_wellord1(X1)=>v1_wellord1(X2))))))), inference(split_equiv,[status(thm)],[c_0_3])).
fof(c_0_7, plain, ![X8, X9, X10]:(~v1_relat_1(X8)|(~v1_relat_1(X9)|(~v1_relat_1(X10)|~v1_funct_1(X10)|(~r3_wellord1(X8,X9,X10)|epred1_2(X9,X8))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
fof(c_0_8, negated_conjecture, (v1_relat_1(esk1_0)&(v1_relat_1(esk2_0)&((v1_relat_1(esk3_0)&v1_funct_1(esk3_0))&((v2_wellord1(esk1_0)&r3_wellord1(esk1_0,esk2_0,esk3_0))&~v2_wellord1(esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
fof(c_0_9, plain, ![X11, X12]:((~v1_relat_2(X11)|v1_relat_2(X12)|~epred1_2(X12,X11))&((~v8_relat_2(X11)|v8_relat_2(X12)|~epred1_2(X12,X11))&((~v6_relat_2(X11)|v6_relat_2(X12)|~epred1_2(X12,X11))&((~v4_relat_2(X11)|v4_relat_2(X12)|~epred1_2(X12,X11))&(~v1_wellord1(X11)|v1_wellord1(X12)|~epred1_2(X12,X11)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
cnf(c_0_10, plain, (epred1_2(X2,X1)|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X3)|~v1_funct_1(X3)|~r3_wellord1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, negated_conjecture, (r3_wellord1(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16, plain, (v1_wellord1(X2)|~v1_wellord1(X1)|~epred1_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, negated_conjecture, (epred1_2(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12]), c_0_13]), c_0_14]), c_0_15])])).
fof(c_0_18, plain, ![X7]:(((v1_relat_2(X7)|~v2_wellord1(X7)|~v1_relat_1(X7))&((v8_relat_2(X7)|~v2_wellord1(X7)|~v1_relat_1(X7))&((v4_relat_2(X7)|~v2_wellord1(X7)|~v1_relat_1(X7))&((v6_relat_2(X7)|~v2_wellord1(X7)|~v1_relat_1(X7))&(v1_wellord1(X7)|~v2_wellord1(X7)|~v1_relat_1(X7))))))&(~v1_relat_2(X7)|(~v8_relat_2(X7)|(~v4_relat_2(X7)|(~v6_relat_2(X7)|~v1_wellord1(X7))))|v2_wellord1(X7)|~v1_relat_1(X7))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_wellord1])])])).
cnf(c_0_19, plain, (v6_relat_2(X2)|~v6_relat_2(X1)|~epred1_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_20, plain, (v1_wellord1(esk2_0)|~v1_wellord1(esk1_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_21, plain, (v1_wellord1(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, negated_conjecture, (v2_wellord1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_23, plain, (v6_relat_2(esk2_0)|~v6_relat_2(esk1_0)), inference(spm,[status(thm)],[c_0_19, c_0_17])).
cnf(c_0_24, plain, (v6_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (v1_relat_2(X2)|~v1_relat_2(X1)|~epred1_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_26, plain, (v2_wellord1(X1)|~v1_relat_2(X1)|~v8_relat_2(X1)|~v4_relat_2(X1)|~v6_relat_2(X1)|~v1_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, plain, (v1_wellord1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_15])])).
cnf(c_0_28, plain, (v6_relat_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_22]), c_0_15])])).
cnf(c_0_29, negated_conjecture, (~v2_wellord1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_30, plain, (v1_relat_2(esk2_0)|~v1_relat_2(esk1_0)), inference(spm,[status(thm)],[c_0_25, c_0_17])).
cnf(c_0_31, plain, (v1_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_32, plain, (v8_relat_2(X2)|~v8_relat_2(X1)|~epred1_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_33, plain, (~v4_relat_2(esk2_0)|~v8_relat_2(esk2_0)|~v1_relat_2(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_14])]), c_0_29])).
cnf(c_0_34, plain, (v1_relat_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_22]), c_0_15])])).
cnf(c_0_35, plain, (v8_relat_2(esk2_0)|~v8_relat_2(esk1_0)), inference(spm,[status(thm)],[c_0_32, c_0_17])).
cnf(c_0_36, plain, (v8_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_37, plain, (~v4_relat_2(esk2_0)|~v8_relat_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_34])])).
cnf(c_0_38, plain, (v8_relat_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_22]), c_0_15])])).
cnf(c_0_39, plain, (v4_relat_2(X2)|~v4_relat_2(X1)|~epred1_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_40, plain, (~v4_relat_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_38])])).
cnf(c_0_41, plain, (~v4_relat_2(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_17]), c_0_40])).
cnf(c_0_42, plain, (v4_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_43, plain, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_22]), c_0_15])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 44
# Proof object clause steps            : 33
# Proof object formula steps           : 11
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 29
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 19
# Processed clauses                    : 54
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 51
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 6
# Generated clauses                    : 16
# ...of the previous two non-trivial   : 16
# Contextual simplify-reflections      : 0
# Paramodulations                      : 16
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 26
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 12
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 25
# Clause-clause subsumption calls (NU) : 15
# Rec. Clause-clause subsumption calls : 11
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 27
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1272

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.000 s
# Total time               : 0.009 s
# Maximum resident set size: 28588 pages
