# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C47_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(free_g1_orders_2, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=>![X3, X4]:(g1_orders_2(X1,X2)=g1_orders_2(X3,X4)=>(X1=X3&X2=X4))), file('data/MPTP/proofs/295c3503180b4f19b0414f8e17a5eccd.E_input', free_g1_orders_2)).
fof(dt_k1_yellow_1, axiom, ![X1]:(v1_relat_2(k1_yellow_1(X1))&(v4_relat_2(k1_yellow_1(X1))&(v8_relat_2(k1_yellow_1(X1))&(v1_partfun1(k1_yellow_1(X1),X1)&m1_subset_1(k1_yellow_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))), file('data/MPTP/proofs/295c3503180b4f19b0414f8e17a5eccd.E_input', dt_k1_yellow_1)).
fof(fc5_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&(v3_orders_2(k2_yellow_1(X1))&(v4_orders_2(k2_yellow_1(X1))&v5_orders_2(k2_yellow_1(X1))))), file('data/MPTP/proofs/295c3503180b4f19b0414f8e17a5eccd.E_input', fc5_yellow_1)).
fof(d1_yellow_1, axiom, ![X1]:k2_yellow_1(X1)=g1_orders_2(X1,k1_yellow_1(X1)), file('data/MPTP/proofs/295c3503180b4f19b0414f8e17a5eccd.E_input', d1_yellow_1)).
fof(dt_k2_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&l1_orders_2(k2_yellow_1(X1))), file('data/MPTP/proofs/295c3503180b4f19b0414f8e17a5eccd.E_input', dt_k2_yellow_1)).
fof(d4_tarski, axiom, ![X1, X2]:(X2=k3_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X3,X4)&r2_hidden(X4,X1)))), file('data/MPTP/proofs/295c3503180b4f19b0414f8e17a5eccd.E_input', d4_tarski)).
fof(abstractness_v1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>(v1_orders_2(X1)=>X1=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))), file('data/MPTP/proofs/295c3503180b4f19b0414f8e17a5eccd.E_input', abstractness_v1_orders_2)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/295c3503180b4f19b0414f8e17a5eccd.E_input', d9_lattice3)).
fof(d5_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>(v2_yellow_0(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&r2_lattice3(X1,u1_struct_0(X1),X2)))), file('data/MPTP/proofs/295c3503180b4f19b0414f8e17a5eccd.E_input', d5_yellow_0)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/295c3503180b4f19b0414f8e17a5eccd.E_input', t1_subset)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/295c3503180b4f19b0414f8e17a5eccd.E_input', t2_subset)).
fof(t15_yellow_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>(v2_yellow_0(k2_yellow_1(X1))=>r2_hidden(k3_tarski(X1),X1))), file('data/MPTP/proofs/295c3503180b4f19b0414f8e17a5eccd.E_input', t15_yellow_1)).
fof(t3_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r3_orders_2(k2_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3))))), file('data/MPTP/proofs/295c3503180b4f19b0414f8e17a5eccd.E_input', t3_yellow_1)).
fof(fc6_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>(~(v2_struct_0(k2_yellow_1(X1)))&v1_orders_2(k2_yellow_1(X1)))), file('data/MPTP/proofs/295c3503180b4f19b0414f8e17a5eccd.E_input', fc6_yellow_1)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/295c3503180b4f19b0414f8e17a5eccd.E_input', redefinition_r3_orders_2)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/295c3503180b4f19b0414f8e17a5eccd.E_input', d3_tarski)).
fof(c_0_16, plain, ![X15, X16, X17, X18]:((X15=X17|g1_orders_2(X15,X16)!=g1_orders_2(X17,X18)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X15,X15))))&(X16=X18|g1_orders_2(X15,X16)!=g1_orders_2(X17,X18)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X15,X15))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_orders_2])])])])).
fof(c_0_17, plain, ![X55]:(v1_relat_2(k1_yellow_1(X55))&(v4_relat_2(k1_yellow_1(X55))&(v8_relat_2(k1_yellow_1(X55))&(v1_partfun1(k1_yellow_1(X55),X55)&m1_subset_1(k1_yellow_1(X55),k1_zfmisc_1(k2_zfmisc_1(X55,X55))))))), inference(variable_rename,[status(thm)],[dt_k1_yellow_1])).
fof(c_0_18, plain, ![X19]:(v1_orders_2(k2_yellow_1(X19))&(v3_orders_2(k2_yellow_1(X19))&(v4_orders_2(k2_yellow_1(X19))&v5_orders_2(k2_yellow_1(X19))))), inference(variable_rename,[status(thm)],[fc5_yellow_1])).
fof(c_0_19, plain, ![X14]:k2_yellow_1(X14)=g1_orders_2(X14,k1_yellow_1(X14)), inference(variable_rename,[status(thm)],[d1_yellow_1])).
fof(c_0_20, plain, ![X29]:(v1_orders_2(k2_yellow_1(X29))&l1_orders_2(k2_yellow_1(X29))), inference(variable_rename,[status(thm)],[dt_k2_yellow_1])).
fof(c_0_21, plain, ![X44, X45, X46, X48, X49, X50, X51, X53]:((((r2_hidden(X46,esk5_3(X44,X45,X46))|~r2_hidden(X46,X45)|X45!=k3_tarski(X44))&(r2_hidden(esk5_3(X44,X45,X46),X44)|~r2_hidden(X46,X45)|X45!=k3_tarski(X44)))&(~r2_hidden(X48,X49)|~r2_hidden(X49,X44)|r2_hidden(X48,X45)|X45!=k3_tarski(X44)))&((~r2_hidden(esk6_2(X50,X51),X51)|(~r2_hidden(esk6_2(X50,X51),X53)|~r2_hidden(X53,X50))|X51=k3_tarski(X50))&((r2_hidden(esk6_2(X50,X51),esk7_2(X50,X51))|r2_hidden(esk6_2(X50,X51),X51)|X51=k3_tarski(X50))&(r2_hidden(esk7_2(X50,X51),X50)|r2_hidden(esk6_2(X50,X51),X51)|X51=k3_tarski(X50))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_tarski])])])])])])).
cnf(c_0_22, plain, (X1=X2|g1_orders_2(X1,X3)!=g1_orders_2(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (m1_subset_1(k1_yellow_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_24, plain, ![X10]:(~l1_orders_2(X10)|(~v1_orders_2(X10)|X10=g1_orders_2(u1_struct_0(X10),u1_orders_2(X10)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_orders_2])])).
cnf(c_0_25, plain, (v1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, plain, (k2_yellow_1(X1)=g1_orders_2(X1,k1_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (l1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_28, plain, ![X24, X25, X26, X27]:((~r2_lattice3(X24,X25,X26)|(~m1_subset_1(X27,u1_struct_0(X24))|(~r2_hidden(X27,X25)|r1_orders_2(X24,X27,X26)))|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))&((m1_subset_1(esk2_3(X24,X25,X26),u1_struct_0(X24))|r2_lattice3(X24,X25,X26)|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))&((r2_hidden(esk2_3(X24,X25,X26),X25)|r2_lattice3(X24,X25,X26)|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))&(~r1_orders_2(X24,esk2_3(X24,X25,X26),X26)|r2_lattice3(X24,X25,X26)|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
fof(c_0_29, plain, ![X32, X34]:(((m1_subset_1(esk3_1(X32),u1_struct_0(X32))|~v2_yellow_0(X32)|~l1_orders_2(X32))&(r2_lattice3(X32,u1_struct_0(X32),esk3_1(X32))|~v2_yellow_0(X32)|~l1_orders_2(X32)))&(~m1_subset_1(X34,u1_struct_0(X32))|~r2_lattice3(X32,u1_struct_0(X32),X34)|v2_yellow_0(X32)|~l1_orders_2(X32))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_yellow_0])])])])])).
fof(c_0_30, plain, ![X30, X31]:(~r2_hidden(X30,X31)|m1_subset_1(X30,X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_31, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X2,X3)|X4!=k3_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_32, plain, ![X8, X9]:(~m1_subset_1(X8,X9)|(v1_xboole_0(X9)|r2_hidden(X8,X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_33, plain, (X1=X2|g1_orders_2(X1,k1_yellow_1(X1))!=g1_orders_2(X2,X3)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_34, plain, (X1=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~l1_orders_2(X1)|~v1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_35, plain, (v1_orders_2(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_36, plain, (l1_orders_2(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_27, c_0_26])).
cnf(c_0_37, plain, (r1_orders_2(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, plain, (r2_lattice3(X1,u1_struct_0(X1),esk3_1(X1))|~v2_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40, plain, (m1_subset_1(esk3_1(X1),u1_struct_0(X1))|~v2_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_41, plain, (r2_hidden(X1,k3_tarski(X2))|~r2_hidden(X3,X2)|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_31])).
cnf(c_0_42, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43, plain, (u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1)))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34])]), c_0_35]), c_0_36])])).
fof(c_0_44, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>(v2_yellow_0(k2_yellow_1(X1))=>r2_hidden(k3_tarski(X1),X1)))), inference(assume_negation,[status(cth)],[t15_yellow_1])).
cnf(c_0_45, plain, (r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk2_3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_46, plain, (r1_orders_2(X1,X2,esk3_1(X1))|~l1_orders_2(X1)|~r2_hidden(X2,u1_struct_0(X1))|~v2_yellow_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39]), c_0_40])).
cnf(c_0_47, plain, (r2_hidden(X1,esk5_3(X2,X3,X1))|~r2_hidden(X1,X3)|X3!=k3_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_48, plain, (r2_hidden(X1,k3_tarski(X2))|v1_xboole_0(X2)|~m1_subset_1(X3,X2)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_49, plain, (m1_subset_1(esk3_1(g1_orders_2(X1,k1_yellow_1(X1))),X1)|~v2_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_43]), c_0_36])])).
fof(c_0_50, negated_conjecture, (~v1_xboole_0(esk1_0)&(v2_yellow_0(k2_yellow_1(esk1_0))&~r2_hidden(k3_tarski(esk1_0),esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_44])])])])).
fof(c_0_51, plain, ![X35, X36, X37]:((~r3_orders_2(k2_yellow_1(X35),X36,X37)|r1_tarski(X36,X37)|~m1_subset_1(X37,u1_struct_0(k2_yellow_1(X35)))|~m1_subset_1(X36,u1_struct_0(k2_yellow_1(X35)))|v1_xboole_0(X35))&(~r1_tarski(X36,X37)|r3_orders_2(k2_yellow_1(X35),X36,X37)|~m1_subset_1(X37,u1_struct_0(k2_yellow_1(X35)))|~m1_subset_1(X36,u1_struct_0(k2_yellow_1(X35)))|v1_xboole_0(X35))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_yellow_1])])])])])).
cnf(c_0_52, plain, (r2_lattice3(X1,X2,esk3_1(X1))|~l1_orders_2(X1)|~r2_hidden(esk2_3(X1,X2,esk3_1(X1)),u1_struct_0(X1))|~v2_yellow_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_40])).
cnf(c_0_53, plain, (X2=k3_tarski(X1)|~r2_hidden(esk6_2(X1,X2),X2)|~r2_hidden(esk6_2(X1,X2),X3)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_54, plain, (r2_hidden(X1,esk5_3(X2,k3_tarski(X2),X1))|~r2_hidden(X1,k3_tarski(X2))), inference(er,[status(thm)],[c_0_47])).
cnf(c_0_55, plain, (r2_hidden(esk5_3(X1,X2,X3),X1)|~r2_hidden(X3,X2)|X2!=k3_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_56, plain, (r2_hidden(X1,k3_tarski(X2))|v1_xboole_0(X2)|~r2_hidden(X1,esk3_1(g1_orders_2(X2,k1_yellow_1(X2))))|~v2_yellow_0(g1_orders_2(X2,k1_yellow_1(X2)))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_57, plain, (r2_hidden(esk7_2(X1,X2),X1)|r2_hidden(esk6_2(X1,X2),X2)|X2=k3_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_58, negated_conjecture, (v2_yellow_0(k2_yellow_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_59, plain, (r1_tarski(X2,X3)|v1_xboole_0(X1)|~r3_orders_2(k2_yellow_1(X1),X2,X3)|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_51])).
fof(c_0_60, plain, ![X20]:((~v2_struct_0(k2_yellow_1(X20))|v1_xboole_0(X20))&(v1_orders_2(k2_yellow_1(X20))|v1_xboole_0(X20))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc6_yellow_1])])])])).
cnf(c_0_61, plain, (r2_lattice3(X1,X2,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(esk2_3(X1,X2,esk3_1(X1)),u1_struct_0(X1))|~v2_yellow_0(X1)), inference(spm,[status(thm)],[c_0_52, c_0_42])).
cnf(c_0_62, plain, (m1_subset_1(esk2_3(X1,X2,X3),u1_struct_0(X1))|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_63, plain, (X1=k3_tarski(X2)|~r2_hidden(esk5_3(X3,k3_tarski(X3),esk6_2(X2,X1)),X2)|~r2_hidden(esk6_2(X2,X1),k3_tarski(X3))|~r2_hidden(esk6_2(X2,X1),X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_64, plain, (r2_hidden(esk5_3(X1,k3_tarski(X1),X2),X1)|~r2_hidden(X2,k3_tarski(X1))), inference(er,[status(thm)],[c_0_55])).
cnf(c_0_65, plain, (esk3_1(g1_orders_2(X1,k1_yellow_1(X1)))=k3_tarski(X2)|r2_hidden(esk6_2(X2,esk3_1(g1_orders_2(X1,k1_yellow_1(X1)))),k3_tarski(X1))|r2_hidden(esk7_2(X2,esk3_1(g1_orders_2(X1,k1_yellow_1(X1)))),X2)|v1_xboole_0(X1)|~v2_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_66, negated_conjecture, (v2_yellow_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))), inference(rw,[status(thm)],[c_0_58, c_0_26])).
cnf(c_0_67, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_68, plain, (v1_xboole_0(X1)|r1_tarski(X2,X3)|~r3_orders_2(g1_orders_2(X1,k1_yellow_1(X1)),X2,X3)|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_26]), c_0_26]), c_0_26])).
fof(c_0_69, plain, ![X21, X22, X23]:((~r3_orders_2(X21,X22,X23)|r1_orders_2(X21,X22,X23)|(v2_struct_0(X21)|(~v3_orders_2(X21)|~l1_orders_2(X21))|(~m1_subset_1(X22,u1_struct_0(X21))|~m1_subset_1(X23,u1_struct_0(X21)))))&(~r1_orders_2(X21,X22,X23)|r3_orders_2(X21,X22,X23)|(v2_struct_0(X21)|(~v3_orders_2(X21)|~l1_orders_2(X21))|(~m1_subset_1(X22,u1_struct_0(X21))|~m1_subset_1(X23,u1_struct_0(X21)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_70, plain, (v3_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_71, plain, (v1_xboole_0(X1)|~v2_struct_0(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_72, plain, (r2_lattice3(X1,X2,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)|~v2_yellow_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_40])).
cnf(c_0_73, plain, (X1=k3_tarski(X2)|~r2_hidden(esk6_2(X2,X1),k3_tarski(X2))|~r2_hidden(esk6_2(X2,X1),X1)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_74, negated_conjecture, (esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))=k3_tarski(X1)|r2_hidden(esk6_2(X1,esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))),k3_tarski(esk1_0))|r2_hidden(esk7_2(X1,esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67])).
cnf(c_0_75, plain, (r1_tarski(X1,X2)|v1_xboole_0(X3)|~r3_orders_2(g1_orders_2(X3,k1_yellow_1(X3)),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_43]), c_0_43])).
cnf(c_0_76, plain, (r3_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_77, plain, (v3_orders_2(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_70, c_0_26])).
cnf(c_0_78, plain, (v1_xboole_0(X1)|~v2_struct_0(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_71, c_0_26])).
cnf(c_0_79, plain, (r1_orders_2(X1,X2,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X2,X3)|~v2_yellow_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_72]), c_0_40])).
cnf(c_0_80, negated_conjecture, (esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))=k3_tarski(esk1_0)|r2_hidden(esk7_2(esk1_0,esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))),esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_57])).
fof(c_0_81, plain, ![X38, X39, X40, X41, X42]:((~r1_tarski(X38,X39)|(~r2_hidden(X40,X38)|r2_hidden(X40,X39)))&((r2_hidden(esk4_2(X41,X42),X41)|r1_tarski(X41,X42))&(~r2_hidden(esk4_2(X41,X42),X42)|r1_tarski(X41,X42)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_82, plain, (r1_tarski(X1,X2)|v1_xboole_0(X3)|~r1_orders_2(g1_orders_2(X3,k1_yellow_1(X3)),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77]), c_0_36]), c_0_43]), c_0_43])]), c_0_78])).
cnf(c_0_83, negated_conjecture, (esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))=k3_tarski(esk1_0)|r1_orders_2(X1,esk7_2(esk1_0,esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))),esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(esk7_2(esk1_0,esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))),u1_struct_0(X1))|~v2_yellow_0(X1)), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_84, plain, (X1=k3_tarski(X2)|r2_hidden(esk7_2(X2,X1),X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_57]), c_0_57])).
cnf(c_0_85, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_86, negated_conjecture, (esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))=k3_tarski(esk1_0)|r1_tarski(esk7_2(esk1_0,esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))),esk3_1(g1_orders_2(X1,k1_yellow_1(X1))))|v1_xboole_0(X1)|~m1_subset_1(esk7_2(esk1_0,esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))),X1)|~v2_yellow_0(g1_orders_2(X1,k1_yellow_1(X1)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_43]), c_0_36]), c_0_43])]), c_0_49])).
cnf(c_0_87, plain, (X1=k3_tarski(X2)|r2_hidden(X3,k3_tarski(X2))|~r2_hidden(X3,esk7_2(X2,X1))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_41, c_0_84])).
cnf(c_0_88, plain, (r2_hidden(esk6_2(X1,X2),esk7_2(X1,X2))|r2_hidden(esk6_2(X1,X2),X2)|X2=k3_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_89, negated_conjecture, (esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))=k3_tarski(esk1_0)|r2_hidden(X1,esk3_1(g1_orders_2(X2,k1_yellow_1(X2))))|v1_xboole_0(X2)|~m1_subset_1(esk7_2(esk1_0,esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))),X2)|~r2_hidden(X1,esk7_2(esk1_0,esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))))|~v2_yellow_0(g1_orders_2(X2,k1_yellow_1(X2)))), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_90, negated_conjecture, (esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))=k3_tarski(esk1_0)|m1_subset_1(esk7_2(esk1_0,esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))),esk1_0)), inference(spm,[status(thm)],[c_0_39, c_0_80])).
cnf(c_0_91, plain, (X1=k3_tarski(X2)|r2_hidden(esk6_2(X2,X1),k3_tarski(X2))|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_41])).
cnf(c_0_92, negated_conjecture, (esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))=k3_tarski(esk1_0)|r2_hidden(X1,esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))|~r2_hidden(X1,esk7_2(esk1_0,esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_66])]), c_0_67])).
cnf(c_0_93, plain, (X1=k3_tarski(X2)|~r2_hidden(esk6_2(X2,X1),X1)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_73, c_0_91])).
cnf(c_0_94, negated_conjecture, (esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))=k3_tarski(esk1_0)|r2_hidden(esk6_2(esk1_0,esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))),esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))), inference(spm,[status(thm)],[c_0_92, c_0_88])).
cnf(c_0_95, negated_conjecture, (esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))=k3_tarski(esk1_0)|~r2_hidden(esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))),esk1_0)), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_96, negated_conjecture, (esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))=k3_tarski(esk1_0)|~m1_subset_1(esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))),esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_42]), c_0_67])).
cnf(c_0_97, negated_conjecture, (~r2_hidden(k3_tarski(esk1_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_98, negated_conjecture, (esk3_1(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))=k3_tarski(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_49]), c_0_66])])).
cnf(c_0_99, negated_conjecture, (~m1_subset_1(k3_tarski(esk1_0),esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_42]), c_0_67])).
cnf(c_0_100, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_98]), c_0_43]), c_0_36]), c_0_66])]), c_0_99]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 101
# Proof object clause steps            : 68
# Proof object formula steps           : 33
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 31
# Proof object simplifying inferences  : 51
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 46
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 45
# Processed clauses                    : 477
# ...of these trivial                  : 3
# ...subsumed                          : 116
# ...remaining for further processing  : 358
# Other redundant clauses eliminated   : 14
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 79
# Generated clauses                    : 3230
# ...of the previous two non-trivial   : 3203
# Contextual simplify-reflections      : 13
# Paramodulations                      : 3198
# Factorizations                       : 9
# Equation resolutions                 : 23
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 227
#    Positive orientable unit clauses  : 15
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 209
# Current number of unprocessed clauses: 2767
# ...number of literals in the above   : 15894
# Current number of archived formulas  : 0
# Current number of archived clauses   : 127
# Clause-clause subsumption calls (NU) : 24875
# Rec. Clause-clause subsumption calls : 6192
# Non-unit clause-clause subsumptions  : 133
# Unit Clause-clause subsumption calls : 3
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 103782

# -------------------------------------------------
# User time                : 0.069 s
# System time              : 0.000 s
# Total time               : 0.069 s
# Maximum resident set size: 28672 pages
