# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_orders_2(X1,X2,X3)<=>(r1_orders_2(X1,X2,X3)&~(X2=X3)))))), file('data/MPTP/proofs/29c118d1a4ab4ef1b5696084f64e9894.E_input', d10_orders_2)).
fof(fraenkel_a_2_1_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&(v3_orders_2(X2)&(v4_orders_2(X2)&(v5_orders_2(X2)&l1_orders_2(X2)))))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r2_hidden(X1,a_2_1_orders_2(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r2_hidden(X5,X3)=>r2_orders_2(X2,X4,X5))))))), file('data/MPTP/proofs/29c118d1a4ab4ef1b5696084f64e9894.E_input', fraenkel_a_2_1_orders_2)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/29c118d1a4ab4ef1b5696084f64e9894.E_input', t4_subset)).
fof(dt_k2_orders_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k2_orders_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/29c118d1a4ab4ef1b5696084f64e9894.E_input', dt_k2_orders_2)).
fof(d13_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k2_orders_2(X1,X2)=a_2_1_orders_2(X1,X2))), file('data/MPTP/proofs/29c118d1a4ab4ef1b5696084f64e9894.E_input', d13_orders_2)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/29c118d1a4ab4ef1b5696084f64e9894.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/29c118d1a4ab4ef1b5696084f64e9894.E_input', t2_xboole_1)).
fof(t1_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&r1_xboole_0(X2,X1))))), file('data/MPTP/proofs/29c118d1a4ab4ef1b5696084f64e9894.E_input', t1_mcart_1)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/29c118d1a4ab4ef1b5696084f64e9894.E_input', l3_subset_1)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/29c118d1a4ab4ef1b5696084f64e9894.E_input', t4_subset_1)).
fof(t46_orders_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>k2_orders_2(X1,k2_struct_0(X1))=k1_xboole_0), file('data/MPTP/proofs/29c118d1a4ab4ef1b5696084f64e9894.E_input', t46_orders_2)).
fof(dt_k2_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/29c118d1a4ab4ef1b5696084f64e9894.E_input', dt_k2_struct_0)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/29c118d1a4ab4ef1b5696084f64e9894.E_input', d3_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/29c118d1a4ab4ef1b5696084f64e9894.E_input', dt_l1_orders_2)).
fof(c_0_14, plain, ![X9, X10, X11]:(((r1_orders_2(X9,X10,X11)|~r2_orders_2(X9,X10,X11)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|~l1_orders_2(X9))&(X10!=X11|~r2_orders_2(X9,X10,X11)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|~l1_orders_2(X9)))&(~r1_orders_2(X9,X10,X11)|X10=X11|r2_orders_2(X9,X10,X11)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|~l1_orders_2(X9))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_orders_2])])])])])).
fof(c_0_15, plain, ![X17, X18, X19, X21, X22]:(((m1_subset_1(esk3_3(X17,X18,X19),u1_struct_0(X18))|~r2_hidden(X17,a_2_1_orders_2(X18,X19))|(v2_struct_0(X18)|(~v3_orders_2(X18)|(~v4_orders_2(X18)|(~v5_orders_2(X18)|~l1_orders_2(X18))))|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))&((X17=esk3_3(X17,X18,X19)|~r2_hidden(X17,a_2_1_orders_2(X18,X19))|(v2_struct_0(X18)|(~v3_orders_2(X18)|(~v4_orders_2(X18)|(~v5_orders_2(X18)|~l1_orders_2(X18))))|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))&(~m1_subset_1(X21,u1_struct_0(X18))|(~r2_hidden(X21,X19)|r2_orders_2(X18,esk3_3(X17,X18,X19),X21))|~r2_hidden(X17,a_2_1_orders_2(X18,X19))|(v2_struct_0(X18)|(~v3_orders_2(X18)|(~v4_orders_2(X18)|(~v5_orders_2(X18)|~l1_orders_2(X18))))|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))))&((m1_subset_1(esk4_4(X17,X18,X19,X22),u1_struct_0(X18))|X17!=X22|~m1_subset_1(X22,u1_struct_0(X18))|r2_hidden(X17,a_2_1_orders_2(X18,X19))|(v2_struct_0(X18)|(~v3_orders_2(X18)|(~v4_orders_2(X18)|(~v5_orders_2(X18)|~l1_orders_2(X18))))|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))&((r2_hidden(esk4_4(X17,X18,X19,X22),X19)|X17!=X22|~m1_subset_1(X22,u1_struct_0(X18))|r2_hidden(X17,a_2_1_orders_2(X18,X19))|(v2_struct_0(X18)|(~v3_orders_2(X18)|(~v4_orders_2(X18)|(~v5_orders_2(X18)|~l1_orders_2(X18))))|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))&(~r2_orders_2(X18,X22,esk4_4(X17,X18,X19,X22))|X17!=X22|~m1_subset_1(X22,u1_struct_0(X18))|r2_hidden(X17,a_2_1_orders_2(X18,X19))|(v2_struct_0(X18)|(~v3_orders_2(X18)|(~v4_orders_2(X18)|(~v5_orders_2(X18)|~l1_orders_2(X18))))|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_1_orders_2])])])])])])).
fof(c_0_16, plain, ![X12, X13, X14]:(~r2_hidden(X12,X13)|~m1_subset_1(X13,k1_zfmisc_1(X14))|m1_subset_1(X12,X14)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_17, plain, ![X7, X8]:(v2_struct_0(X7)|(~v3_orders_2(X7)|(~v4_orders_2(X7)|(~v5_orders_2(X7)|~l1_orders_2(X7))))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|m1_subset_1(k2_orders_2(X7,X8),k1_zfmisc_1(u1_struct_0(X7)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_orders_2])])])).
fof(c_0_18, plain, ![X25, X26]:(v2_struct_0(X25)|(~v3_orders_2(X25)|(~v4_orders_2(X25)|(~v5_orders_2(X25)|~l1_orders_2(X25))))|(~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|k2_orders_2(X25,X26)=a_2_1_orders_2(X25,X26))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_orders_2])])])])).
fof(c_0_19, plain, ![X34, X35]:(~r2_hidden(X34,X35)|~r1_tarski(X35,X34)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_20, plain, ![X31]:r1_tarski(k1_xboole_0,X31), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_21, plain, (X1!=X2|~r2_orders_2(X3,X1,X2)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (r2_orders_2(X2,esk3_3(X4,X2,X3),X1)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r2_hidden(X4,a_2_1_orders_2(X2,X3))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_24, plain, ![X15]:((r2_hidden(esk2_1(X15),X15)|X15=k1_xboole_0)&(r1_xboole_0(esk2_1(X15),X15)|X15=k1_xboole_0)), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t1_mcart_1])])])])])).
fof(c_0_25, plain, ![X27, X28, X29]:(~m1_subset_1(X28,k1_zfmisc_1(X27))|(~r2_hidden(X29,X28)|r2_hidden(X29,X27))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
cnf(c_0_26, plain, (v2_struct_0(X1)|m1_subset_1(k2_orders_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27, plain, (v2_struct_0(X1)|k2_orders_2(X1,X2)=a_2_1_orders_2(X1,X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30, plain, (r2_hidden(esk4_4(X1,X2,X3,X4),X3)|r2_hidden(X1,a_2_1_orders_2(X2,X3))|v2_struct_0(X2)|X1!=X4|~m1_subset_1(X4,u1_struct_0(X2))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_31, plain, ![X30]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X30)), inference(variable_rename,[status(thm)],[t4_subset_1])).
fof(c_0_32, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>k2_orders_2(X1,k2_struct_0(X1))=k1_xboole_0)), inference(assume_negation,[status(cth)],[t46_orders_2])).
cnf(c_0_33, plain, (~r2_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_21])).
cnf(c_0_34, plain, (r2_orders_2(X1,esk3_3(X2,X1,X3),X4)|v2_struct_0(X1)|~r2_hidden(X2,a_2_1_orders_2(X1,X3))|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_35, plain, (X1=esk3_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_orders_2(X2,X3))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_36, plain, (r2_hidden(esk2_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_37, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_38, plain, (m1_subset_1(a_2_1_orders_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_39, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_40, plain, (r2_hidden(esk4_4(X1,X2,X3,X1),X3)|r2_hidden(X1,a_2_1_orders_2(X2,X3))|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(er,[status(thm)],[c_0_30])).
cnf(c_0_41, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_42, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&k2_orders_2(esk1_0,k2_struct_0(esk1_0))!=k1_xboole_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_32])])])])).
cnf(c_0_43, plain, (v2_struct_0(X1)|~r2_hidden(esk3_3(X2,X1,X3),X3)|~r2_hidden(X2,a_2_1_orders_2(X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_23])).
cnf(c_0_44, plain, (esk3_3(esk2_1(a_2_1_orders_2(X1,X2)),X1,X2)=esk2_1(a_2_1_orders_2(X1,X2))|a_2_1_orders_2(X1,X2)=k1_xboole_0|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_45, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_orders_2(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_46, plain, (r2_hidden(X1,a_2_1_orders_2(X2,k1_xboole_0))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])])).
cnf(c_0_47, negated_conjecture, (k2_orders_2(esk1_0,k2_struct_0(esk1_0))!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_48, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_49, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_50, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_51, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_52, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_53, plain, ![X32]:(~l1_struct_0(X32)|m1_subset_1(k2_struct_0(X32),k1_zfmisc_1(u1_struct_0(X32)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_struct_0])])).
cnf(c_0_54, plain, (a_2_1_orders_2(X1,X2)=k1_xboole_0|v2_struct_0(X1)|~r2_hidden(esk2_1(a_2_1_orders_2(X1,X2)),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_36])).
cnf(c_0_55, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_41])])).
cnf(c_0_56, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_orders_2(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(spm,[status(thm)],[c_0_23, c_0_38])).
fof(c_0_57, plain, ![X24]:(~l1_struct_0(X24)|k2_struct_0(X24)=u1_struct_0(X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_58, negated_conjecture, (a_2_1_orders_2(esk1_0,k2_struct_0(esk1_0))!=k1_xboole_0|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_27]), c_0_48]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])).
cnf(c_0_59, plain, (m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_60, plain, (a_2_1_orders_2(X1,u1_struct_0(X2))=k1_xboole_0|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk2_1(a_2_1_orders_2(X1,u1_struct_0(X2))),u1_struct_0(X2))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~l1_orders_2(X2)|~v5_orders_2(X1)|~v5_orders_2(X2)|~v4_orders_2(X1)|~v4_orders_2(X2)|~v3_orders_2(X1)|~v3_orders_2(X2)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_61, plain, (a_2_1_orders_2(X1,X2)=k1_xboole_0|m1_subset_1(esk2_1(a_2_1_orders_2(X1,X2)),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_56, c_0_36])).
cnf(c_0_62, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
fof(c_0_63, plain, ![X33]:(~l1_orders_2(X33)|l1_struct_0(X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_64, negated_conjecture, (a_2_1_orders_2(esk1_0,k2_struct_0(esk1_0))!=k1_xboole_0|~l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_65, plain, (a_2_1_orders_2(X1,u1_struct_0(X1))=k1_xboole_0|v2_struct_0(X1)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_66, plain, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_59, c_0_62])).
cnf(c_0_67, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_68, negated_conjecture, (a_2_1_orders_2(esk1_0,u1_struct_0(esk1_0))!=k1_xboole_0|~l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_64, c_0_62])).
cnf(c_0_69, plain, (a_2_1_orders_2(X1,u1_struct_0(X1))=k1_xboole_0|v2_struct_0(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67])).
cnf(c_0_70, negated_conjecture, (~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_48]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])).
cnf(c_0_71, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_67]), c_0_48])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 72
# Proof object clause steps            : 43
# Proof object formula steps           : 29
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 24
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 109
# ...of these trivial                  : 0
# ...subsumed                          : 16
# ...remaining for further processing  : 93
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 0
# Generated clauses                    : 91
# ...of the previous two non-trivial   : 82
# Contextual simplify-reflections      : 4
# Paramodulations                      : 87
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 60
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 50
# Current number of unprocessed clauses: 27
# ...number of literals in the above   : 329
# Current number of archived formulas  : 0
# Current number of archived clauses   : 29
# Clause-clause subsumption calls (NU) : 3077
# Rec. Clause-clause subsumption calls : 152
# Non-unit clause-clause subsumptions  : 14
# Unit Clause-clause subsumption calls : 48
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5931

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.004 s
# Total time               : 0.014 s
# Maximum resident set size: 28784 pages
