# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t17_subset_1, conjecture, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(X1))=>(![X5]:(m1_subset_1(X5,X1)=>(r2_hidden(X5,X2)<=>(r2_hidden(X5,X3)&~(r2_hidden(X5,X4)))))=>X2=k7_subset_1(X1,X3,X4))))), file('data/MPTP/proofs/2a7c8ee3213a49a69b6d759ac75eeaac.E_input', t17_subset_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/2a7c8ee3213a49a69b6d759ac75eeaac.E_input', d2_subset_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/2a7c8ee3213a49a69b6d759ac75eeaac.E_input', d1_xboole_0)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/2a7c8ee3213a49a69b6d759ac75eeaac.E_input', l3_subset_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/2a7c8ee3213a49a69b6d759ac75eeaac.E_input', d5_xboole_0)).
fof(redefinition_k7_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k7_subset_1(X1,X2,X3)=k4_xboole_0(X2,X3)), file('data/MPTP/proofs/2a7c8ee3213a49a69b6d759ac75eeaac.E_input', redefinition_k7_subset_1)).
fof(c_0_6, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(X1))=>(![X5]:(m1_subset_1(X5,X1)=>(r2_hidden(X5,X2)<=>(r2_hidden(X5,X3)&~(r2_hidden(X5,X4)))))=>X2=k7_subset_1(X1,X3,X4)))))), inference(assume_negation,[status(cth)],[t17_subset_1])).
fof(c_0_7, plain, ![X14, X15]:(((~m1_subset_1(X15,X14)|r2_hidden(X15,X14)|v1_xboole_0(X14))&(~r2_hidden(X15,X14)|m1_subset_1(X15,X14)|v1_xboole_0(X14)))&((~m1_subset_1(X15,X14)|v1_xboole_0(X15)|~v1_xboole_0(X14))&(~v1_xboole_0(X15)|m1_subset_1(X15,X14)|~v1_xboole_0(X14)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_8, plain, ![X28, X29, X30]:((~v1_xboole_0(X28)|~r2_hidden(X29,X28))&(r2_hidden(esk6_1(X30),X30)|v1_xboole_0(X30))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_9, plain, ![X11, X12, X13]:(~m1_subset_1(X12,k1_zfmisc_1(X11))|(~r2_hidden(X13,X12)|r2_hidden(X13,X11))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
fof(c_0_10, negated_conjecture, ![X10]:(m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk4_0,k1_zfmisc_1(esk1_0))&((((r2_hidden(X10,esk3_0)|~r2_hidden(X10,esk2_0)|~m1_subset_1(X10,esk1_0))&(~r2_hidden(X10,esk4_0)|~r2_hidden(X10,esk2_0)|~m1_subset_1(X10,esk1_0)))&(~r2_hidden(X10,esk3_0)|r2_hidden(X10,esk4_0)|r2_hidden(X10,esk2_0)|~m1_subset_1(X10,esk1_0)))&esk2_0!=k7_subset_1(esk1_0,esk3_0,esk4_0))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_6])])])])])])).
cnf(c_0_11, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, plain, ![X19, X20, X21, X22, X23, X24, X25, X26]:((((r2_hidden(X22,X19)|~r2_hidden(X22,X21)|X21!=k4_xboole_0(X19,X20))&(~r2_hidden(X22,X20)|~r2_hidden(X22,X21)|X21!=k4_xboole_0(X19,X20)))&(~r2_hidden(X23,X19)|r2_hidden(X23,X20)|r2_hidden(X23,X21)|X21!=k4_xboole_0(X19,X20)))&((~r2_hidden(esk5_3(X24,X25,X26),X26)|(~r2_hidden(esk5_3(X24,X25,X26),X24)|r2_hidden(esk5_3(X24,X25,X26),X25))|X26=k4_xboole_0(X24,X25))&((r2_hidden(esk5_3(X24,X25,X26),X24)|r2_hidden(esk5_3(X24,X25,X26),X26)|X26=k4_xboole_0(X24,X25))&(~r2_hidden(esk5_3(X24,X25,X26),X25)|r2_hidden(esk5_3(X24,X25,X26),X26)|X26=k4_xboole_0(X24,X25))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_16, negated_conjecture, (r2_hidden(X1,esk4_0)|r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_18, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_19, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_20, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_21, negated_conjecture, (~r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_22, plain, (r2_hidden(esk5_3(X1,X2,X3),X1)|r2_hidden(esk5_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (r2_hidden(esk5_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk5_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, negated_conjecture, (r2_hidden(X1,esk4_0)|r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])).
fof(c_0_25, plain, ![X16, X17, X18]:(~m1_subset_1(X17,k1_zfmisc_1(X16))|k7_subset_1(X16,X17,X18)=k4_xboole_0(X17,X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_subset_1])])).
cnf(c_0_26, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_19, c_0_17])).
cnf(c_0_27, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_13, c_0_20])).
cnf(c_0_28, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_29, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_30, negated_conjecture, (~r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_21, c_0_17])).
cnf(c_0_31, plain, (k4_xboole_0(X1,X2)=X1|r2_hidden(esk5_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_22])).
cnf(c_0_32, plain, (r2_hidden(esk5_3(X1,X2,X3),X2)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk5_3(X1,X2,X3),X3)|~r2_hidden(esk5_3(X1,X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_33, negated_conjecture, (X1=k4_xboole_0(X2,esk4_0)|r2_hidden(esk5_3(X2,esk4_0,X1),esk2_0)|r2_hidden(esk5_3(X2,esk4_0,X1),X1)|~r2_hidden(esk5_3(X2,esk4_0,X1),esk3_0)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_34, negated_conjecture, (esk2_0!=k7_subset_1(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_35, plain, (k7_subset_1(X2,X1,X3)=k4_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_36, negated_conjecture, (X1=k4_xboole_0(X2,esk3_0)|r2_hidden(esk5_3(X2,esk3_0,X1),X1)|~r2_hidden(esk5_3(X2,esk3_0,X1),esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_26]), c_0_27])).
cnf(c_0_37, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_28])).
cnf(c_0_38, plain, (r2_hidden(esk6_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_39, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_29])).
cnf(c_0_40, negated_conjecture, (k4_xboole_0(esk4_0,X1)=esk4_0|~r2_hidden(esk5_3(esk4_0,X1,esk4_0),esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_27])).
cnf(c_0_41, plain, (k4_xboole_0(X1,X2)=X1|r2_hidden(esk5_3(X1,X2,X1),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_31]), c_0_31])).
cnf(c_0_42, negated_conjecture, (X1=k4_xboole_0(esk3_0,esk4_0)|r2_hidden(esk5_3(esk3_0,esk4_0,X1),esk2_0)|r2_hidden(esk5_3(esk3_0,esk4_0,X1),X1)), inference(spm,[status(thm)],[c_0_33, c_0_22])).
cnf(c_0_43, negated_conjecture, (k4_xboole_0(esk3_0,esk4_0)!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_14])])).
cnf(c_0_44, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_45, negated_conjecture, (X1=k4_xboole_0(esk2_0,esk3_0)|r2_hidden(esk5_3(esk2_0,esk3_0,X1),X1)), inference(spm,[status(thm)],[c_0_36, c_0_22])).
cnf(c_0_46, plain, (v1_xboole_0(k4_xboole_0(X1,X2))|~r2_hidden(esk6_1(k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_47, plain, (v1_xboole_0(k4_xboole_0(X1,X2))|r2_hidden(esk6_1(k4_xboole_0(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_39, c_0_38])).
cnf(c_0_48, negated_conjecture, (k4_xboole_0(esk4_0,esk2_0)=esk4_0), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_49, negated_conjecture, (r2_hidden(esk5_3(esk3_0,esk4_0,esk2_0),esk2_0)), inference(sr,[status(thm)],[inference(ef,[status(thm)],[c_0_42]), c_0_43])).
cnf(c_0_50, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_51, negated_conjecture, (X1=k4_xboole_0(esk2_0,esk3_0)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_12, c_0_45])).
cnf(c_0_52, plain, (v1_xboole_0(k4_xboole_0(X1,X1))), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_53, negated_conjecture, (~r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_37, c_0_48])).
cnf(c_0_54, negated_conjecture, (r2_hidden(esk5_3(esk3_0,esk4_0,esk2_0),esk4_0)|~r2_hidden(esk5_3(esk3_0,esk4_0,esk2_0),esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_49]), c_0_43])).
cnf(c_0_55, plain, (r2_hidden(X1,X2)|~v1_xboole_0(k4_xboole_0(X3,X2))|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_12, c_0_50])).
cnf(c_0_56, negated_conjecture, (k4_xboole_0(X1,X1)=k4_xboole_0(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_57, negated_conjecture, (~r2_hidden(esk5_3(esk3_0,esk4_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_49])])).
cnf(c_0_58, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_52])])).
cnf(c_0_59, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_49])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 60
# Proof object clause steps            : 47
# Proof object formula steps           : 13
# Proof object conjectures             : 29
# Proof object clause conjectures      : 26
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 26
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 21
# Processed clauses                    : 5470
# ...of these trivial                  : 39
# ...subsumed                          : 4678
# ...remaining for further processing  : 753
# Other redundant clauses eliminated   : 14
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 73
# Backward-rewritten                   : 50
# Generated clauses                    : 33568
# ...of the previous two non-trivial   : 30475
# Contextual simplify-reflections      : 30
# Paramodulations                      : 33372
# Factorizations                       : 178
# Equation resolutions                 : 14
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 602
#    Positive orientable unit clauses  : 42
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 13
#    Non-unit-clauses                  : 545
# Current number of unprocessed clauses: 24659
# ...number of literals in the above   : 101185
# Current number of archived formulas  : 0
# Current number of archived clauses   : 148
# Clause-clause subsumption calls (NU) : 49006
# Rec. Clause-clause subsumption calls : 32334
# Non-unit clause-clause subsumptions  : 2349
# Unit Clause-clause subsumption calls : 891
# Rewrite failures with RHS unbound    : 51
# BW rewrite match attempts            : 150
# BW rewrite match successes           : 30
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 605491

# -------------------------------------------------
# User time                : 0.388 s
# System time              : 0.003 s
# Total time               : 0.392 s
# Maximum resident set size: 28680 pages
