# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(dt_k8_mcart_1, axiom, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>m1_subset_1(k8_mcart_1(X1,X2,X3,X4,X5),X1)), file('data/MPTP/proofs/2b1c333e93b44c60a0c137a885e82c06.E_input', dt_k8_mcart_1)).
fof(t53_mcart_1, axiom, ![X1, X2, X3, X4]:k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4), file('data/MPTP/proofs/2b1c333e93b44c60a0c137a885e82c06.E_input', t53_mcart_1)).
fof(t61_mcart_1, axiom, ![X1, X2, X3, X4]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&~(![X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>(k8_mcart_1(X1,X2,X3,X4,X5)=k1_mcart_1(k1_mcart_1(k1_mcart_1(X5)))&(k9_mcart_1(X1,X2,X3,X4,X5)=k2_mcart_1(k1_mcart_1(k1_mcart_1(X5)))&(k10_mcart_1(X1,X2,X3,X4,X5)=k2_mcart_1(k1_mcart_1(X5))&k11_mcart_1(X1,X2,X3,X4,X5)=k2_mcart_1(X5))))))))))), file('data/MPTP/proofs/2b1c333e93b44c60a0c137a885e82c06.E_input', t61_mcart_1)).
fof(t80_mcart_1, conjecture, ![X1, X2, X3, X4, X5, X6]:(m1_subset_1(X6,k4_zfmisc_1(X1,X2,X3,X4))=>(![X7]:(m1_subset_1(X7,X1)=>![X8]:(m1_subset_1(X8,X2)=>![X9]:(m1_subset_1(X9,X3)=>![X10]:(m1_subset_1(X10,X4)=>(X6=k4_mcart_1(X7,X8,X9,X10)=>X5=X8)))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|(X4=k1_xboole_0|X5=k9_mcart_1(X1,X2,X3,X4,X6))))))), file('data/MPTP/proofs/2b1c333e93b44c60a0c137a885e82c06.E_input', t80_mcart_1)).
fof(t31_mcart_1, axiom, ![X1, X2, X3, X4]:k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4), file('data/MPTP/proofs/2b1c333e93b44c60a0c137a885e82c06.E_input', t31_mcart_1)).
fof(t60_mcart_1, axiom, ![X1, X2, X3, X4]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&~(![X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>X5=k4_mcart_1(k8_mcart_1(X1,X2,X3,X4,X5),k9_mcart_1(X1,X2,X3,X4,X5),k10_mcart_1(X1,X2,X3,X4,X5),k11_mcart_1(X1,X2,X3,X4,X5))))))))), file('data/MPTP/proofs/2b1c333e93b44c60a0c137a885e82c06.E_input', t60_mcart_1)).
fof(dt_k11_mcart_1, axiom, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>m1_subset_1(k11_mcart_1(X1,X2,X3,X4,X5),X4)), file('data/MPTP/proofs/2b1c333e93b44c60a0c137a885e82c06.E_input', dt_k11_mcart_1)).
fof(dt_k10_mcart_1, axiom, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>m1_subset_1(k10_mcart_1(X1,X2,X3,X4,X5),X3)), file('data/MPTP/proofs/2b1c333e93b44c60a0c137a885e82c06.E_input', dt_k10_mcart_1)).
fof(dt_k9_mcart_1, axiom, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>m1_subset_1(k9_mcart_1(X1,X2,X3,X4,X5),X2)), file('data/MPTP/proofs/2b1c333e93b44c60a0c137a885e82c06.E_input', dt_k9_mcart_1)).
fof(c_0_9, plain, ![X11, X12, X13, X14, X15]:(~m1_subset_1(X15,k4_zfmisc_1(X11,X12,X13,X14))|m1_subset_1(k8_mcart_1(X11,X12,X13,X14,X15),X11)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_mcart_1])])).
fof(c_0_10, plain, ![X16, X17, X18, X19]:k4_zfmisc_1(X16,X17,X18,X19)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X16,X17),X18),X19), inference(variable_rename,[status(thm)],[t53_mcart_1])).
fof(c_0_11, plain, ![X20, X21, X22, X23, X24]:((k8_mcart_1(X20,X21,X22,X23,X24)=k1_mcart_1(k1_mcart_1(k1_mcart_1(X24)))|~m1_subset_1(X24,k4_zfmisc_1(X20,X21,X22,X23))|X23=k1_xboole_0|X22=k1_xboole_0|X21=k1_xboole_0|X20=k1_xboole_0)&((k9_mcart_1(X20,X21,X22,X23,X24)=k2_mcart_1(k1_mcart_1(k1_mcart_1(X24)))|~m1_subset_1(X24,k4_zfmisc_1(X20,X21,X22,X23))|X23=k1_xboole_0|X22=k1_xboole_0|X21=k1_xboole_0|X20=k1_xboole_0)&((k10_mcart_1(X20,X21,X22,X23,X24)=k2_mcart_1(k1_mcart_1(X24))|~m1_subset_1(X24,k4_zfmisc_1(X20,X21,X22,X23))|X23=k1_xboole_0|X22=k1_xboole_0|X21=k1_xboole_0|X20=k1_xboole_0)&(k11_mcart_1(X20,X21,X22,X23,X24)=k2_mcart_1(X24)|~m1_subset_1(X24,k4_zfmisc_1(X20,X21,X22,X23))|X23=k1_xboole_0|X22=k1_xboole_0|X21=k1_xboole_0|X20=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t61_mcart_1])])])])])).
fof(c_0_12, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6]:(m1_subset_1(X6,k4_zfmisc_1(X1,X2,X3,X4))=>(![X7]:(m1_subset_1(X7,X1)=>![X8]:(m1_subset_1(X8,X2)=>![X9]:(m1_subset_1(X9,X3)=>![X10]:(m1_subset_1(X10,X4)=>(X6=k4_mcart_1(X7,X8,X9,X10)=>X5=X8)))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|(X4=k1_xboole_0|X5=k9_mcart_1(X1,X2,X3,X4,X6)))))))), inference(assume_negation,[status(cth)],[t80_mcart_1])).
cnf(c_0_13, plain, (m1_subset_1(k8_mcart_1(X2,X3,X4,X5,X1),X2)|~m1_subset_1(X1,k4_zfmisc_1(X2,X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, plain, (k8_mcart_1(X1,X2,X3,X4,X5)=k1_mcart_1(k1_mcart_1(k1_mcart_1(X5)))|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_16, negated_conjecture, ![X31, X32, X33, X34]:(m1_subset_1(esk6_0,k4_zfmisc_1(esk1_0,esk2_0,esk3_0,esk4_0))&((~m1_subset_1(X31,esk1_0)|(~m1_subset_1(X32,esk2_0)|(~m1_subset_1(X33,esk3_0)|(~m1_subset_1(X34,esk4_0)|(esk6_0!=k4_mcart_1(X31,X32,X33,X34)|esk5_0=X32)))))&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&(esk4_0!=k1_xboole_0&esk5_0!=k9_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).
cnf(c_0_17, plain, (m1_subset_1(k8_mcart_1(X2,X3,X4,X5,X1),X2)|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4),X5))), inference(rw,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_18, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k8_mcart_1(X1,X2,X3,X4,X5)=k1_mcart_1(k1_mcart_1(k1_mcart_1(X5)))|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[c_0_15, c_0_14])).
cnf(c_0_19, negated_conjecture, (m1_subset_1(esk6_0,k4_zfmisc_1(esk1_0,esk2_0,esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_20, plain, ![X35, X36, X37, X38]:k4_mcart_1(X35,X36,X37,X38)=k4_tarski(k4_tarski(k4_tarski(X35,X36),X37),X38), inference(variable_rename,[status(thm)],[t31_mcart_1])).
fof(c_0_21, plain, ![X39, X40, X41, X42, X43]:(X39=k1_xboole_0|(X40=k1_xboole_0|(X41=k1_xboole_0|(X42=k1_xboole_0|(~m1_subset_1(X43,k4_zfmisc_1(X39,X40,X41,X42))|X43=k4_mcart_1(k8_mcart_1(X39,X40,X41,X42,X43),k9_mcart_1(X39,X40,X41,X42,X43),k10_mcart_1(X39,X40,X41,X42,X43),k11_mcart_1(X39,X40,X41,X42,X43))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t60_mcart_1])])])])).
cnf(c_0_22, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|m1_subset_1(k1_mcart_1(k1_mcart_1(k1_mcart_1(X5))),X1)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_23, negated_conjecture, (m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),esk4_0))), inference(rw,[status(thm)],[c_0_19, c_0_14])).
cnf(c_0_24, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_27, negated_conjecture, (esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28, negated_conjecture, (esk5_0=X2|~m1_subset_1(X1,esk1_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X4,esk4_0)|esk6_0!=k4_mcart_1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_29, plain, (k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|X5=k4_mcart_1(k8_mcart_1(X1,X2,X3,X4,X5),k9_mcart_1(X1,X2,X3,X4,X5),k10_mcart_1(X1,X2,X3,X4,X5),k11_mcart_1(X1,X2,X3,X4,X5))|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(k1_mcart_1(k1_mcart_1(k1_mcart_1(esk6_0))),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27])).
fof(c_0_32, plain, ![X44, X45, X46, X47, X48]:(~m1_subset_1(X48,k4_zfmisc_1(X44,X45,X46,X47))|m1_subset_1(k11_mcart_1(X44,X45,X46,X47,X48),X47)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k11_mcart_1])])).
cnf(c_0_33, negated_conjecture, (esk5_0=X2|esk6_0!=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)|~m1_subset_1(X4,esk4_0)|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk1_0)), inference(rw,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X5=k4_tarski(k4_tarski(k4_tarski(k8_mcart_1(X1,X2,X3,X4,X5),k9_mcart_1(X1,X2,X3,X4,X5)),k10_mcart_1(X1,X2,X3,X4,X5)),k11_mcart_1(X1,X2,X3,X4,X5))|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_14]), c_0_29])).
cnf(c_0_35, negated_conjecture, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|m1_subset_1(k8_mcart_1(X1,X2,X3,X4,esk6_0),esk1_0)|~m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(spm,[status(thm)],[c_0_31, c_0_18])).
cnf(c_0_36, plain, (m1_subset_1(k11_mcart_1(X2,X3,X4,X5,X1),X5)|~m1_subset_1(X1,k4_zfmisc_1(X2,X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_37, plain, ![X49, X50, X51, X52, X53]:(~m1_subset_1(X53,k4_zfmisc_1(X49,X50,X51,X52))|m1_subset_1(k10_mcart_1(X49,X50,X51,X52,X53),X51)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k10_mcart_1])])).
cnf(c_0_38, negated_conjecture, (k9_mcart_1(X1,X2,X3,X4,esk6_0)=esk5_0|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|~m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))|~m1_subset_1(k11_mcart_1(X1,X2,X3,X4,esk6_0),esk4_0)|~m1_subset_1(k10_mcart_1(X1,X2,X3,X4,esk6_0),esk3_0)|~m1_subset_1(k9_mcart_1(X1,X2,X3,X4,esk6_0),esk2_0)), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34])]), c_0_35])).
cnf(c_0_39, plain, (m1_subset_1(k11_mcart_1(X2,X3,X4,X5,X1),X5)|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4),X5))), inference(rw,[status(thm)],[c_0_36, c_0_14])).
cnf(c_0_40, plain, (m1_subset_1(k10_mcart_1(X2,X3,X4,X5,X1),X4)|~m1_subset_1(X1,k4_zfmisc_1(X2,X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_41, plain, ![X54, X55, X56, X57, X58]:(~m1_subset_1(X58,k4_zfmisc_1(X54,X55,X56,X57))|m1_subset_1(k9_mcart_1(X54,X55,X56,X57,X58),X55)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k9_mcart_1])])).
cnf(c_0_42, negated_conjecture, (k9_mcart_1(X1,X2,X3,esk4_0,esk6_0)=esk5_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),esk4_0))|~m1_subset_1(k10_mcart_1(X1,X2,X3,esk4_0,esk6_0),esk3_0)|~m1_subset_1(k9_mcart_1(X1,X2,X3,esk4_0,esk6_0),esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_27])).
cnf(c_0_43, plain, (m1_subset_1(k10_mcart_1(X2,X3,X4,X5,X1),X4)|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4),X5))), inference(rw,[status(thm)],[c_0_40, c_0_14])).
cnf(c_0_44, plain, (m1_subset_1(k9_mcart_1(X2,X3,X4,X5,X1),X3)|~m1_subset_1(X1,k4_zfmisc_1(X2,X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_45, negated_conjecture, (k9_mcart_1(X1,X2,esk3_0,esk4_0,esk6_0)=esk5_0|X1=k1_xboole_0|X2=k1_xboole_0|~m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),esk3_0),esk4_0))|~m1_subset_1(k9_mcart_1(X1,X2,esk3_0,esk4_0,esk6_0),esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_26])).
cnf(c_0_46, plain, (m1_subset_1(k9_mcart_1(X2,X3,X4,X5,X1),X3)|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4),X5))), inference(rw,[status(thm)],[c_0_44, c_0_14])).
cnf(c_0_47, negated_conjecture, (k9_mcart_1(X1,esk2_0,esk3_0,esk4_0,esk6_0)=esk5_0|X1=k1_xboole_0|~m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,esk2_0),esk3_0),esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_25])).
cnf(c_0_48, negated_conjecture, (esk5_0!=k9_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_49, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_23]), c_0_48]), c_0_24]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 50
# Proof object clause steps            : 31
# Proof object formula steps           : 19
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 20
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 16
# Processed clauses                    : 53
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 53
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 105
# ...of the previous two non-trivial   : 102
# Contextual simplify-reflections      : 1
# Paramodulations                      : 104
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 37
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 27
# Current number of unprocessed clauses: 81
# ...number of literals in the above   : 885
# Current number of archived formulas  : 0
# Current number of archived clauses   : 18
# Clause-clause subsumption calls (NU) : 133
# Rec. Clause-clause subsumption calls : 44
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 18
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5543

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.003 s
# Total time               : 0.019 s
# Maximum resident set size: 28744 pages
