# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C45_F1_PI_AE_Q4_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.006 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t47_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/2b8873055cb3461dacf9de551821a32d.E_input', t47_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/2b8873055cb3461dacf9de551821a32d.E_input', t48_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/2b8873055cb3461dacf9de551821a32d.E_input', commutativity_k3_xboole_0)).
fof(l98_xboole_1, axiom, ![X1, X2]:k5_xboole_0(X1,X2)=k4_xboole_0(k2_xboole_0(X1,X2),k3_xboole_0(X1,X2)), file('data/MPTP/proofs/2b8873055cb3461dacf9de551821a32d.E_input', l98_xboole_1)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/2b8873055cb3461dacf9de551821a32d.E_input', t98_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/2b8873055cb3461dacf9de551821a32d.E_input', t3_boole)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/2b8873055cb3461dacf9de551821a32d.E_input', t41_xboole_1)).
fof(t5_boole, axiom, ![X1]:k5_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/2b8873055cb3461dacf9de551821a32d.E_input', t5_boole)).
fof(t100_xboole_1, conjecture, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/2b8873055cb3461dacf9de551821a32d.E_input', t100_xboole_1)).
fof(c_0_9, plain, ![X4, X5]:k4_xboole_0(X4,k3_xboole_0(X4,X5))=k4_xboole_0(X4,X5), inference(variable_rename,[status(thm)],[t47_xboole_1])).
fof(c_0_10, plain, ![X6, X7]:k4_xboole_0(X6,k4_xboole_0(X6,X7))=k3_xboole_0(X6,X7), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_11, plain, ![X8, X9]:k3_xboole_0(X8,X9)=k3_xboole_0(X9,X8), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_12, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_15, plain, ![X10, X11]:k5_xboole_0(X10,X11)=k4_xboole_0(k2_xboole_0(X10,X11),k3_xboole_0(X10,X11)), inference(variable_rename,[status(thm)],[l98_xboole_1])).
fof(c_0_16, plain, ![X12, X13]:k2_xboole_0(X12,X13)=k5_xboole_0(X12,k4_xboole_0(X13,X12)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_17, plain, ![X14]:k4_xboole_0(X14,k1_xboole_0)=X14, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_18, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X1,X2)))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_19, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_13]), c_0_13])).
cnf(c_0_20, plain, (k5_xboole_0(X1,X2)=k4_xboole_0(k2_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, plain, ![X15, X16, X17]:k4_xboole_0(k4_xboole_0(X15,X16),X17)=k4_xboole_0(X15,k2_xboole_0(X16,X17)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
cnf(c_0_23, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (k4_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X1)))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25, plain, (k5_xboole_0(X1,X2)=k4_xboole_0(k5_xboole_0(X1,k4_xboole_0(X2,X1)),k4_xboole_0(X1,k4_xboole_0(X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_21]), c_0_13])).
cnf(c_0_26, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27, plain, (k4_xboole_0(k1_xboole_0,k4_xboole_0(k1_xboole_0,X1))=k4_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_19, c_0_23])).
cnf(c_0_28, plain, (k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,X1))=k4_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_24, c_0_23])).
cnf(c_0_29, plain, (k4_xboole_0(k5_xboole_0(X1,k4_xboole_0(X2,X1)),k4_xboole_0(X2,k4_xboole_0(X2,X1)))=k5_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_25, c_0_19])).
cnf(c_0_30, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X1))=X1), inference(spm,[status(thm)],[c_0_18, c_0_23])).
cnf(c_0_31, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k5_xboole_0(X2,k4_xboole_0(X3,X2)))), inference(rw,[status(thm)],[c_0_26, c_0_21])).
cnf(c_0_32, plain, (k4_xboole_0(k4_xboole_0(X1,X1),X1)=k4_xboole_0(k1_xboole_0,k4_xboole_0(k1_xboole_0,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_24]), c_0_27]), c_0_28])).
cnf(c_0_33, plain, (k4_xboole_0(k5_xboole_0(k4_xboole_0(X1,X1),X1),k4_xboole_0(X1,X1))=k5_xboole_0(k4_xboole_0(X1,X1),X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_34, plain, (k4_xboole_0(X1,k5_xboole_0(k4_xboole_0(X2,X2),X2))=k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X2,X2)),X2)), inference(spm,[status(thm)],[c_0_31, c_0_30])).
cnf(c_0_35, plain, (k4_xboole_0(k4_xboole_0(k1_xboole_0,k4_xboole_0(k1_xboole_0,X1)),X1)=k4_xboole_0(k1_xboole_0,k4_xboole_0(k1_xboole_0,X1))), inference(spm,[status(thm)],[c_0_32, c_0_27])).
cnf(c_0_36, plain, (k4_xboole_0(k4_xboole_0(X1,X1),k4_xboole_0(k1_xboole_0,k4_xboole_0(k1_xboole_0,X1)))=k4_xboole_0(X1,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_28]), c_0_23])).
cnf(c_0_37, plain, (k4_xboole_0(k5_xboole_0(k4_xboole_0(X1,X1),X1),X1)=k4_xboole_0(X1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_33]), c_0_34]), c_0_33]), c_0_34]), c_0_27]), c_0_28]), c_0_35]), c_0_36])).
cnf(c_0_38, plain, (k5_xboole_0(k4_xboole_0(X1,X1),X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_37]), c_0_33]), c_0_34]), c_0_30]), c_0_30])).
cnf(c_0_39, plain, (k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X2,X2)),X2)=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_34, c_0_38])).
fof(c_0_40, plain, ![X18]:k5_xboole_0(X18,k1_xboole_0)=X18, inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_41, plain, (k4_xboole_0(k4_xboole_0(k1_xboole_0,X1),X1)=k4_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_39, c_0_28])).
cnf(c_0_42, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_43, plain, (k4_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(k1_xboole_0,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_41]), c_0_31]), c_0_23])).
cnf(c_0_44, plain, (k4_xboole_0(k5_xboole_0(X1,k4_xboole_0(k1_xboole_0,X1)),k4_xboole_0(X1,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_23]), c_0_42])).
cnf(c_0_45, plain, (k4_xboole_0(X1,k4_xboole_0(k4_xboole_0(X1,X2),X3))=k4_xboole_0(k5_xboole_0(X2,k4_xboole_0(X3,X2)),k4_xboole_0(k5_xboole_0(X2,k4_xboole_0(X3,X2)),X1))), inference(spm,[status(thm)],[c_0_19, c_0_31])).
cnf(c_0_46, plain, (k4_xboole_0(k5_xboole_0(X1,k4_xboole_0(X1,X1)),X1)=k5_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_25, c_0_30])).
cnf(c_0_47, plain, (k4_xboole_0(X1,k5_xboole_0(k4_xboole_0(X2,X3),k4_xboole_0(X3,k4_xboole_0(X3,X2))))=k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X2,X3)),X2)), inference(spm,[status(thm)],[c_0_31, c_0_19])).
cnf(c_0_48, plain, (k4_xboole_0(X1,k4_xboole_0(k1_xboole_0,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_28])).
cnf(c_0_49, plain, (k4_xboole_0(k5_xboole_0(X1,k4_xboole_0(X1,X1)),k5_xboole_0(X1,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_32]), c_0_24]), c_0_23])).
cnf(c_0_50, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_41]), c_0_43]), c_0_31]), c_0_23])).
cnf(c_0_51, plain, (k4_xboole_0(k4_xboole_0(k5_xboole_0(X1,X1),X1),X1)=k4_xboole_0(k1_xboole_0,k4_xboole_0(k1_xboole_0,k5_xboole_0(X1,X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_49]), c_0_46]), c_0_31]), c_0_27])).
cnf(c_0_52, plain, (k4_xboole_0(k5_xboole_0(X1,X1),X1)=k5_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_50, c_0_46])).
cnf(c_0_53, plain, (k4_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X2,X2)))=k4_xboole_0(k4_xboole_0(k4_xboole_0(X1,X2),X2),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_46]), c_0_31])).
cnf(c_0_54, plain, (k5_xboole_0(k4_xboole_0(k1_xboole_0,k4_xboole_0(k1_xboole_0,X1)),X1)=X1), inference(spm,[status(thm)],[c_0_38, c_0_27])).
cnf(c_0_55, plain, (k4_xboole_0(k1_xboole_0,k4_xboole_0(k1_xboole_0,k5_xboole_0(X1,X1)))=k5_xboole_0(X1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_50]), c_0_52])).
cnf(c_0_56, plain, (k4_xboole_0(k5_xboole_0(X1,X1),k4_xboole_0(k1_xboole_0,X1))=k5_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_43, c_0_46])).
cnf(c_0_57, plain, (k4_xboole_0(k4_xboole_0(X1,X2),k5_xboole_0(X2,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_52]), c_0_53]), c_0_50]), c_0_50])).
cnf(c_0_58, plain, (k5_xboole_0(k5_xboole_0(X1,X1),k5_xboole_0(X1,X1))=k5_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_59, plain, (k5_xboole_0(X1,X1)=k4_xboole_0(k1_xboole_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_56]), c_0_27]), c_0_55]), c_0_57]), c_0_27]), c_0_18])).
cnf(c_0_60, plain, (k4_xboole_0(k1_xboole_0,k4_xboole_0(k1_xboole_0,X1))=k4_xboole_0(k1_xboole_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_59]), c_0_59]), c_0_59])).
cnf(c_0_61, plain, (k4_xboole_0(k1_xboole_0,k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,X2)))=k4_xboole_0(k4_xboole_0(X1,X2),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_19]), c_0_24]), c_0_27])).
cnf(c_0_62, plain, (k4_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X3,X2))=k4_xboole_0(k4_xboole_0(X1,X2),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_50]), c_0_31])).
cnf(c_0_63, plain, (k4_xboole_0(X1,k5_xboole_0(k5_xboole_0(X2,k4_xboole_0(X3,X2)),k4_xboole_0(k4_xboole_0(X4,X2),X3)))=k4_xboole_0(k4_xboole_0(k4_xboole_0(X1,X2),X3),X4)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_31]), c_0_31])).
cnf(c_0_64, plain, (k4_xboole_0(k1_xboole_0,X1)=k4_xboole_0(X1,X1)), inference(rw,[status(thm)],[c_0_27, c_0_60])).
cnf(c_0_65, plain, (k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k4_xboole_0(k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X1),X2))=k4_xboole_0(k5_xboole_0(X1,k4_xboole_0(X2,X1)),k5_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_25]), c_0_31])).
cnf(c_0_66, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X1)=k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[c_0_61, c_0_60])).
cnf(c_0_67, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,X1)))=k4_xboole_0(k1_xboole_0,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_62]), c_0_27]), c_0_60])).
cnf(c_0_68, plain, (k4_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,k4_xboole_0(X1,X2)))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_50, c_0_18])).
cnf(c_0_69, plain, (k4_xboole_0(k4_xboole_0(k4_xboole_0(X1,X2),X3),X2)=k4_xboole_0(k4_xboole_0(X1,X2),X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_63]), c_0_23])).
cnf(c_0_70, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(k4_xboole_0(X1,X2),X3)))=k4_xboole_0(k4_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_18, c_0_31])).
cnf(c_0_71, plain, (k4_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(k1_xboole_0,X1)))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_62, c_0_48])).
cnf(c_0_72, plain, (k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k4_xboole_0(k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,k4_xboole_0(X1,X2))),X2))=k4_xboole_0(k5_xboole_0(X1,k4_xboole_0(X2,X1)),k5_xboole_0(X1,X2))), inference(rw,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_73, plain, (k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,k4_xboole_0(X1,X2)))=k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_50]), c_0_62]), c_0_66])).
cnf(c_0_74, plain, (k4_xboole_0(k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,X2)),X2)=k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_69, c_0_64])).
cnf(c_0_75, plain, (k4_xboole_0(k4_xboole_0(k5_xboole_0(X1,X2),X1),X2)=k4_xboole_0(k1_xboole_0,k5_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_25]), c_0_31])).
cnf(c_0_76, plain, (k4_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(k1_xboole_0,X1))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_18])).
cnf(c_0_77, plain, (k4_xboole_0(k5_xboole_0(X1,k4_xboole_0(X2,X1)),k5_xboole_0(X1,X2))=k4_xboole_0(X1,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_73]), c_0_74]), c_0_62]), c_0_23])).
cnf(c_0_78, plain, (k4_xboole_0(k1_xboole_0,k4_xboole_0(k5_xboole_0(X1,X2),X1))=k4_xboole_0(k1_xboole_0,k5_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_75]), c_0_76]), c_0_62]), c_0_66])).
cnf(c_0_79, plain, (k4_xboole_0(k1_xboole_0,k5_xboole_0(X1,X2))=k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_77]), c_0_25]), c_0_73])).
cnf(c_0_80, plain, (k5_xboole_0(k4_xboole_0(k1_xboole_0,X1),X1)=X1), inference(rw,[status(thm)],[c_0_54, c_0_60])).
cnf(c_0_81, plain, (k4_xboole_0(k1_xboole_0,k4_xboole_0(k5_xboole_0(X1,X2),X1))=k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_82, plain, (k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,X2))=k4_xboole_0(k1_xboole_0,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_19]), c_0_24]), c_0_73])).
cnf(c_0_83, plain, (k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(k1_xboole_0,X2)),X2)=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_64]), c_0_30]), c_0_80])).
cnf(c_0_84, plain, (k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,k5_xboole_0(X1,X2)))=k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_85, plain, (k4_xboole_0(k1_xboole_0,k4_xboole_0(k4_xboole_0(X1,X1),X2))=k4_xboole_0(k4_xboole_0(X1,X1),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_64]), c_0_60])).
cnf(c_0_86, plain, (k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,k4_xboole_0(X2,X1)))=k4_xboole_0(k1_xboole_0,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_67]), c_0_62]), c_0_23]), c_0_62]), c_0_66]), c_0_60])).
cnf(c_0_87, plain, (k5_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(k1_xboole_0,X1)))=k5_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_71]), c_0_83]), c_0_25])).
cnf(c_0_88, plain, (k4_xboole_0(X1,k4_xboole_0(k1_xboole_0,k4_xboole_0(X2,X1)))=k4_xboole_0(X1,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_62]), c_0_62]), c_0_66])).
cnf(c_0_89, plain, (k4_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X1,X1)))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_62, c_0_30])).
cnf(c_0_90, plain, (k4_xboole_0(k4_xboole_0(X1,X1),X2)=k4_xboole_0(k1_xboole_0,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_31]), c_0_85]), c_0_86])).
cnf(c_0_91, plain, (k5_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,k4_xboole_0(X1,X2)))=k5_xboole_0(k4_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_92, plain, (k5_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X1,X1)))=k5_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_89]), c_0_39]), c_0_25])).
cnf(c_0_93, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X1))=k4_xboole_0(X1,k4_xboole_0(X2,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_90]), c_0_60]), c_0_88])).
cnf(c_0_94, plain, (k4_xboole_0(X1,k5_xboole_0(k4_xboole_0(X2,k4_xboole_0(X2,X3)),k4_xboole_0(X2,X3)))=k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3))),X2)), inference(spm,[status(thm)],[c_0_31, c_0_18])).
cnf(c_0_95, plain, (k4_xboole_0(X1,k5_xboole_0(k4_xboole_0(X2,X3),X3))=k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_31, c_0_91])).
cnf(c_0_96, plain, (k5_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X1,X2)))=k5_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_97, plain, (k4_xboole_0(k4_xboole_0(k1_xboole_0,X1),X2)=k4_xboole_0(k1_xboole_0,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_79]), c_0_86])).
cnf(c_0_98, plain, (k5_xboole_0(X1,k4_xboole_0(k1_xboole_0,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_48]), c_0_41]), c_0_44])).
cnf(c_0_99, plain, (k4_xboole_0(k5_xboole_0(X1,X2),k4_xboole_0(k4_xboole_0(k1_xboole_0,X1),X2))=k5_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_25]), c_0_31])).
cnf(c_0_100, plain, (k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3))),k4_xboole_0(X2,X3))=k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3))),X2)), inference(rw,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_101, plain, (k5_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(k1_xboole_0,X1))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_76]), c_0_97]), c_0_66]), c_0_60]), c_0_98])).
cnf(c_0_102, plain, (k4_xboole_0(k5_xboole_0(X1,X2),k4_xboole_0(k4_xboole_0(X1,X1),X2))=k5_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_99, c_0_64])).
cnf(c_0_103, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(k4_xboole_0(X2,k4_xboole_0(X2,X1)),X3)))=k4_xboole_0(k4_xboole_0(X2,k4_xboole_0(X2,X1)),X3)), inference(spm,[status(thm)],[c_0_70, c_0_19])).
cnf(c_0_104, plain, (k4_xboole_0(k5_xboole_0(X1,X2),k4_xboole_0(X2,X1))=k4_xboole_0(k5_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_100, c_0_29])).
cnf(c_0_105, plain, (k5_xboole_0(k5_xboole_0(X1,X2),k4_xboole_0(k1_xboole_0,k4_xboole_0(X2,X1)))=k5_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_25]), c_0_31]), c_0_97])).
cnf(c_0_106, plain, (k4_xboole_0(X1,k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,X2)))=k4_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_88, c_0_82])).
cnf(c_0_107, plain, (k4_xboole_0(k5_xboole_0(X1,X2),k4_xboole_0(k1_xboole_0,k4_xboole_0(X2,X1)))=k5_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_102, c_0_90])).
cnf(c_0_108, plain, (k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k4_xboole_0(X2,k4_xboole_0(X2,X1)))=k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_24, c_0_24])).
cnf(c_0_109, plain, (k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X2)=k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_93]), c_0_103]), c_0_62]), c_0_66]), c_0_73])).
cnf(c_0_110, plain, (k4_xboole_0(k5_xboole_0(X1,X2),k4_xboole_0(X1,X2))=k4_xboole_0(k5_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_100, c_0_25])).
cnf(c_0_111, plain, (k4_xboole_0(k4_xboole_0(k5_xboole_0(X1,X2),X1),X2)=k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[c_0_75, c_0_79])).
cnf(c_0_112, plain, (k4_xboole_0(k5_xboole_0(X1,X2),k4_xboole_0(X2,X2))=k5_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_97]), c_0_104]), c_0_106]), c_0_107]), c_0_93])).
cnf(c_0_113, plain, (k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k4_xboole_0(X2,k4_xboole_0(X2,X1)))=k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[c_0_108, c_0_109])).
cnf(c_0_114, plain, (k5_xboole_0(X1,k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,X2)))=k5_xboole_0(X1,k4_xboole_0(k1_xboole_0,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_83])).
fof(c_0_115, negated_conjecture, ~(![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(assume_negation,[status(cth)],[t100_xboole_1])).
cnf(c_0_116, plain, (k4_xboole_0(k5_xboole_0(X1,k4_xboole_0(X1,X2)),X1)=k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_110]), c_0_111]), c_0_73]), c_0_50])).
cnf(c_0_117, plain, (k4_xboole_0(k5_xboole_0(X1,X2),k4_xboole_0(k1_xboole_0,X2))=k5_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_30]), c_0_28]), c_0_30])).
cnf(c_0_118, plain, (k4_xboole_0(X1,k5_xboole_0(X2,k4_xboole_0(k1_xboole_0,X3)))=k4_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_66]), c_0_114])).
fof(c_0_119, negated_conjecture, k4_xboole_0(esk1_0,esk2_0)!=k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_115])])])).
cnf(c_0_120, plain, (k4_xboole_0(k5_xboole_0(X1,k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,X2))),k4_xboole_0(X1,X2))=k5_xboole_0(X1,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_19]), c_0_66]), c_0_24])).
cnf(c_0_121, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k5_xboole_0(X1,k4_xboole_0(X1,X2))))=k5_xboole_0(X1,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_116]), c_0_117])).
cnf(c_0_122, plain, (k4_xboole_0(X1,k5_xboole_0(X1,k4_xboole_0(k1_xboole_0,X2)))=k4_xboole_0(k1_xboole_0,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_118]), c_0_66]), c_0_60])).
cnf(c_0_123, negated_conjecture, (k4_xboole_0(esk1_0,esk2_0)!=k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_124, plain, (k4_xboole_0(k5_xboole_0(X1,k4_xboole_0(k1_xboole_0,X2)),k4_xboole_0(X1,X2))=k5_xboole_0(X1,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[c_0_120, c_0_114])).
cnf(c_0_125, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X1))=k5_xboole_0(X1,k4_xboole_0(k1_xboole_0,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_114]), c_0_122]), c_0_106]), c_0_114])).
cnf(c_0_126, plain, (k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X2,X1)),k4_xboole_0(X1,X2))=k4_xboole_0(X1,k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_83, c_0_106])).
cnf(c_0_127, negated_conjecture, (k4_xboole_0(esk1_0,esk2_0)!=k5_xboole_0(esk1_0,k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0)))), inference(rw,[status(thm)],[c_0_123, c_0_13])).
cnf(c_0_128, plain, (k5_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X1,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125]), c_0_126])).
cnf(c_0_129, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_127, c_0_128]), c_0_18])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 130
# Proof object clause steps            : 111
# Proof object formula steps           : 19
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 81
# Proof object simplifying inferences  : 145
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 9
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 7
# Processed clauses                    : 3778
# ...of these trivial                  : 749
# ...subsumed                          : 2604
# ...remaining for further processing  : 425
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 109
# Generated clauses                    : 291208
# ...of the previous two non-trivial   : 248672
# Contextual simplify-reflections      : 0
# Paramodulations                      : 291208
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 316
#    Positive orientable unit clauses  : 280
#    Positive unorientable unit clauses: 36
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 243493
# ...number of literals in the above   : 243493
# Current number of archived formulas  : 0
# Current number of archived clauses   : 111
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 590
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7061
# BW rewrite match successes           : 224
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 10754923

# -------------------------------------------------
# User time                : 4.725 s
# System time              : 0.128 s
# Total time               : 4.853 s
# Maximum resident set size: 264284 pages
