# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/2bc5a5ac9f57488787b6f944b9bdc54d.E_input', d4_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/2bc5a5ac9f57488787b6f944b9bdc54d.E_input', d3_tarski)).
fof(t123_zfmisc_1, axiom, ![X1, X2, X3, X4]:k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)), file('data/MPTP/proofs/2bc5a5ac9f57488787b6f944b9bdc54d.E_input', t123_zfmisc_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/2bc5a5ac9f57488787b6f944b9bdc54d.E_input', idempotence_k3_xboole_0)).
fof(t134_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4)))), file('data/MPTP/proofs/2bc5a5ac9f57488787b6f944b9bdc54d.E_input', t134_zfmisc_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/2bc5a5ac9f57488787b6f944b9bdc54d.E_input', commutativity_k3_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/2bc5a5ac9f57488787b6f944b9bdc54d.E_input', t28_xboole_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/2bc5a5ac9f57488787b6f944b9bdc54d.E_input', d5_xboole_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/2bc5a5ac9f57488787b6f944b9bdc54d.E_input', d1_xboole_0)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/2bc5a5ac9f57488787b6f944b9bdc54d.E_input', t3_boole)).
fof(t7_tarski, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&r2_hidden(X4,X3))))))), file('data/MPTP/proofs/2bc5a5ac9f57488787b6f944b9bdc54d.E_input', t7_tarski)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/2bc5a5ac9f57488787b6f944b9bdc54d.E_input', l13_xboole_0)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/2bc5a5ac9f57488787b6f944b9bdc54d.E_input', t113_zfmisc_1)).
fof(t6_xboole_0, axiom, ![X1, X2]:~((r2_xboole_0(X1,X2)&![X3]:~((r2_hidden(X3,X2)&~(r2_hidden(X3,X1)))))), file('data/MPTP/proofs/2bc5a5ac9f57488787b6f944b9bdc54d.E_input', t6_xboole_0)).
fof(d8_xboole_0, axiom, ![X1, X2]:(r2_xboole_0(X1,X2)<=>(r1_tarski(X1,X2)&~(X1=X2))), file('data/MPTP/proofs/2bc5a5ac9f57488787b6f944b9bdc54d.E_input', d8_xboole_0)).
fof(c_0_15, plain, ![X9, X10, X11, X12, X13, X14, X15, X16]:((((r2_hidden(X12,X9)|~r2_hidden(X12,X11)|X11!=k3_xboole_0(X9,X10))&(r2_hidden(X12,X10)|~r2_hidden(X12,X11)|X11!=k3_xboole_0(X9,X10)))&(~r2_hidden(X13,X9)|~r2_hidden(X13,X10)|r2_hidden(X13,X11)|X11!=k3_xboole_0(X9,X10)))&((~r2_hidden(esk5_3(X14,X15,X16),X16)|(~r2_hidden(esk5_3(X14,X15,X16),X14)|~r2_hidden(esk5_3(X14,X15,X16),X15))|X16=k3_xboole_0(X14,X15))&((r2_hidden(esk5_3(X14,X15,X16),X14)|r2_hidden(esk5_3(X14,X15,X16),X16)|X16=k3_xboole_0(X14,X15))&(r2_hidden(esk5_3(X14,X15,X16),X15)|r2_hidden(esk5_3(X14,X15,X16),X16)|X16=k3_xboole_0(X14,X15))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_16, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_17, plain, ![X32, X33, X34, X35, X36]:((~r1_tarski(X32,X33)|(~r2_hidden(X34,X32)|r2_hidden(X34,X33)))&((r2_hidden(esk8_2(X35,X36),X35)|r1_tarski(X35,X36))&(~r2_hidden(esk8_2(X35,X36),X36)|r1_tarski(X35,X36)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_18, plain, ![X28, X29, X30, X31]:k2_zfmisc_1(k3_xboole_0(X28,X29),k3_xboole_0(X30,X31))=k3_xboole_0(k2_zfmisc_1(X28,X30),k2_zfmisc_1(X29,X31)), inference(variable_rename,[status(thm)],[t123_zfmisc_1])).
fof(c_0_19, plain, ![X47]:k3_xboole_0(X47,X47)=X47, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
fof(c_0_20, negated_conjecture, ~(![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4))))), inference(assume_negation,[status(cth)],[t134_zfmisc_1])).
cnf(c_0_21, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (r2_hidden(esk8_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_23, plain, ![X54, X55]:k3_xboole_0(X54,X55)=k3_xboole_0(X55,X54), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_24, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_26, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k2_zfmisc_1(esk3_0,esk4_0)&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk1_0!=esk3_0|esk2_0!=esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).
cnf(c_0_27, plain, (r1_tarski(X1,X2)|~r2_hidden(esk8_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28, plain, (r1_tarski(k3_xboole_0(X1,X2),X3)|r2_hidden(esk8_2(k3_xboole_0(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_29, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))=k2_zfmisc_1(k3_xboole_0(X1,X3),X2)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k2_zfmisc_1(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_33, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))=k2_zfmisc_1(X1,k3_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_34, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=k3_xboole_0(k2_zfmisc_1(X3,X2),k2_zfmisc_1(X1,X4))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_29]), c_0_24])).
cnf(c_0_35, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(X1,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))=k2_zfmisc_1(k3_xboole_0(X1,esk1_0),esk2_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_36, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X4),k2_zfmisc_1(X3,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_29]), c_0_24])).
cnf(c_0_37, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_32, c_0_29])).
cnf(c_0_38, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk3_0,X1),k2_zfmisc_1(esk1_0,esk4_0))=k2_zfmisc_1(esk1_0,k3_xboole_0(X1,esk2_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_31]), c_0_34])).
fof(c_0_39, plain, ![X18, X19]:(~r1_tarski(X18,X19)|k3_xboole_0(X18,X19)=X18), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_40, plain, ![X38, X39, X40, X41, X42, X43, X44, X45]:((((r2_hidden(X41,X38)|~r2_hidden(X41,X40)|X40!=k4_xboole_0(X38,X39))&(~r2_hidden(X41,X39)|~r2_hidden(X41,X40)|X40!=k4_xboole_0(X38,X39)))&(~r2_hidden(X42,X38)|r2_hidden(X42,X39)|r2_hidden(X42,X40)|X40!=k4_xboole_0(X38,X39)))&((~r2_hidden(esk9_3(X43,X44,X45),X45)|(~r2_hidden(esk9_3(X43,X44,X45),X43)|r2_hidden(esk9_3(X43,X44,X45),X44))|X45=k4_xboole_0(X43,X44))&((r2_hidden(esk9_3(X43,X44,X45),X43)|r2_hidden(esk9_3(X43,X44,X45),X45)|X45=k4_xboole_0(X43,X44))&(~r2_hidden(esk9_3(X43,X44,X45),X44)|r2_hidden(esk9_3(X43,X44,X45),X45)|X45=k4_xboole_0(X43,X44))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_41, plain, ![X48, X49, X50]:((~v1_xboole_0(X48)|~r2_hidden(X49,X48))&(r2_hidden(esk10_1(X50),X50)|v1_xboole_0(X50))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_42, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(X1,esk4_0),k2_zfmisc_1(esk3_0,esk2_0))=k2_zfmisc_1(k3_xboole_0(X1,esk1_0),esk2_0)), inference(rw,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_43, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,k3_xboole_0(X1,esk2_0)),k2_zfmisc_1(esk3_0,X1))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_44, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_45, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_46, plain, (r2_hidden(esk10_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_47, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_48, plain, ![X27]:k4_xboole_0(X27,k1_xboole_0)=X27, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_49, negated_conjecture, (r1_tarski(k2_zfmisc_1(k3_xboole_0(X1,esk1_0),esk2_0),k2_zfmisc_1(X1,esk4_0))), inference(spm,[status(thm)],[c_0_37, c_0_42])).
cnf(c_0_50, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,k3_xboole_0(esk2_0,X1)),k2_zfmisc_1(esk3_0,X1))), inference(spm,[status(thm)],[c_0_43, c_0_29])).
cnf(c_0_51, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_32]), c_0_29])).
cnf(c_0_52, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_45])).
cnf(c_0_53, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|r2_hidden(esk10_1(k3_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_21, c_0_46])).
cnf(c_0_54, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_47])).
cnf(c_0_55, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_56, plain, ![X20, X21, X23]:((r2_hidden(esk6_2(X20,X21),X21)|~r2_hidden(X20,X21))&(~r2_hidden(X23,X21)|~r2_hidden(X23,esk6_2(X20,X21))|~r2_hidden(X20,X21))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tarski])])])])])).
cnf(c_0_57, negated_conjecture, (r1_tarski(k2_zfmisc_1(k3_xboole_0(esk1_0,X1),esk2_0),k2_zfmisc_1(X1,esk4_0))), inference(spm,[status(thm)],[c_0_49, c_0_29])).
cnf(c_0_58, plain, (k3_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_37]), c_0_29])).
cnf(c_0_59, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,k3_xboole_0(X1,esk2_0)),k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk2_0)))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
fof(c_0_60, plain, ![X56]:(~v1_xboole_0(X56)|X56=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_61, plain, (v1_xboole_0(k3_xboole_0(X1,k4_xboole_0(X2,X3)))|~r2_hidden(esk10_1(k3_xboole_0(X1,k4_xboole_0(X2,X3))),X3)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_62, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|r2_hidden(esk10_1(k3_xboole_0(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_53, c_0_29])).
fof(c_0_63, plain, ![X52, X53]:((k2_zfmisc_1(X52,X53)!=k1_xboole_0|(X52=k1_xboole_0|X53=k1_xboole_0))&((X52!=k1_xboole_0|k2_zfmisc_1(X52,X53)=k1_xboole_0)&(X53!=k1_xboole_0|k2_zfmisc_1(X52,X53)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_64, plain, (r1_tarski(k3_xboole_0(X1,k4_xboole_0(X2,X3)),X4)|r2_hidden(esk8_2(k3_xboole_0(X1,k4_xboole_0(X2,X3)),X4),X2)), inference(spm,[status(thm)],[c_0_54, c_0_28])).
cnf(c_0_65, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_52, c_0_55])).
cnf(c_0_66, plain, (r2_hidden(esk6_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_67, negated_conjecture, (r1_tarski(k2_zfmisc_1(k3_xboole_0(esk1_0,X1),esk2_0),k2_zfmisc_1(k3_xboole_0(esk1_0,X1),esk4_0))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_68, negated_conjecture, (k2_zfmisc_1(esk1_0,k3_xboole_0(X1,esk2_0))=k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_59]), c_0_30]), c_0_29]), c_0_24]), c_0_31]), c_0_33])).
cnf(c_0_69, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_70, plain, (v1_xboole_0(k3_xboole_0(X1,k4_xboole_0(X2,X1)))), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_71, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_72, plain, (r1_tarski(k3_xboole_0(X1,k4_xboole_0(X2,X3)),X2)), inference(spm,[status(thm)],[c_0_27, c_0_64])).
cnf(c_0_73, plain, (~r2_hidden(esk6_2(X1,k1_xboole_0),X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_74, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk1_0,esk4_0),k2_zfmisc_1(X1,esk2_0))=k2_zfmisc_1(k3_xboole_0(esk1_0,X1),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_67]), c_0_33]), c_0_29]), c_0_24])).
cnf(c_0_75, negated_conjecture, (k2_zfmisc_1(esk1_0,k3_xboole_0(esk2_0,X1))=k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk4_0))), inference(spm,[status(thm)],[c_0_68, c_0_29])).
cnf(c_0_76, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_77, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_71])).
cnf(c_0_78, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_72, c_0_25])).
cnf(c_0_79, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_73, c_0_66])).
cnf(c_0_80, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(X1,esk3_0),esk4_0)=k2_zfmisc_1(k3_xboole_0(esk1_0,X1),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_74]), c_0_36]), c_0_31]), c_0_30])).
cnf(c_0_81, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_82, negated_conjecture, (k2_zfmisc_1(esk3_0,k3_xboole_0(esk4_0,k4_xboole_0(X1,esk2_0)))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77]), c_0_29])).
cnf(c_0_83, plain, (k3_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_78]), c_0_29])).
cnf(c_0_84, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_79, c_0_22])).
cnf(c_0_85, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk1_0,X1),esk2_0)=k2_zfmisc_1(k3_xboole_0(esk3_0,X1),esk4_0)), inference(spm,[status(thm)],[c_0_80, c_0_29])).
cnf(c_0_86, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_81])).
cnf(c_0_87, negated_conjecture, (r1_tarski(k2_zfmisc_1(k3_xboole_0(X1,esk1_0),esk2_0),k2_zfmisc_1(k3_xboole_0(X1,esk1_0),esk4_0))), inference(spm,[status(thm)],[c_0_57, c_0_51])).
cnf(c_0_88, negated_conjecture, (k2_zfmisc_1(esk3_0,k4_xboole_0(esk4_0,esk2_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_89, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_44, c_0_84])).
cnf(c_0_90, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,X1))=k2_zfmisc_1(esk1_0,k3_xboole_0(esk2_0,X1))), inference(spm,[status(thm)],[c_0_33, c_0_31])).
cnf(c_0_91, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,k4_xboole_0(X4,X2)))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_76]), c_0_77])).
cnf(c_0_92, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk3_0,k4_xboole_0(X1,esk1_0)),esk4_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_76]), c_0_86])).
cnf(c_0_93, plain, (X1=k1_xboole_0|X2=k1_xboole_0|k2_zfmisc_1(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_94, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(X1,esk1_0),esk2_0)=k2_zfmisc_1(k3_xboole_0(X1,esk3_0),esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_87]), c_0_33]), c_0_29]), c_0_24]), c_0_31]), c_0_30])).
cnf(c_0_95, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_96, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk3_0,X1),k4_xboole_0(esk4_0,esk2_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_88]), c_0_89])).
cnf(c_0_97, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_98, negated_conjecture, (k2_zfmisc_1(esk1_0,k3_xboole_0(esk2_0,k4_xboole_0(X1,esk4_0)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_99, negated_conjecture, (k2_zfmisc_1(k4_xboole_0(esk3_0,esk1_0),esk4_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_92, c_0_83])).
cnf(c_0_100, negated_conjecture, (k3_xboole_0(X1,esk1_0)=k1_xboole_0|k2_zfmisc_1(k3_xboole_0(X1,esk3_0),esk4_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_95])).
cnf(c_0_101, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_102, negated_conjecture, (k4_xboole_0(esk4_0,esk2_0)=k1_xboole_0|k3_xboole_0(esk3_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_93, c_0_96])).
cnf(c_0_103, negated_conjecture, (k2_zfmisc_1(esk3_0,esk4_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_31]), c_0_97]), c_0_95])).
cnf(c_0_104, negated_conjecture, (k3_xboole_0(esk2_0,k4_xboole_0(X1,esk4_0))=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_98]), c_0_97])).
cnf(c_0_105, negated_conjecture, (k2_zfmisc_1(k4_xboole_0(esk3_0,esk1_0),k3_xboole_0(esk4_0,X1))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_99]), c_0_89])).
cnf(c_0_106, negated_conjecture, (k3_xboole_0(X1,esk1_0)=k1_xboole_0|k2_zfmisc_1(k3_xboole_0(esk3_0,X1),esk4_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_100, c_0_29])).
fof(c_0_107, plain, ![X24, X25]:((r2_hidden(esk7_2(X24,X25),X25)|~r2_xboole_0(X24,X25))&(~r2_hidden(esk7_2(X24,X25),X24)|~r2_xboole_0(X24,X25))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_xboole_0])])])])])).
cnf(c_0_108, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_101])).
cnf(c_0_109, negated_conjecture, (k4_xboole_0(esk4_0,esk2_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_102]), c_0_86]), c_0_25]), c_0_103])).
cnf(c_0_110, negated_conjecture, (k4_xboole_0(esk2_0,esk4_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_104, c_0_83])).
cnf(c_0_111, negated_conjecture, (k4_xboole_0(esk3_0,esk1_0)=k1_xboole_0|k3_xboole_0(esk4_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_93, c_0_105])).
cnf(c_0_112, negated_conjecture, (k3_xboole_0(esk1_0,k4_xboole_0(X1,esk3_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_76]), c_0_86])]), c_0_29])).
cnf(c_0_113, plain, (~r2_hidden(esk7_2(X1,X2),X1)|~r2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_114, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_79])).
cnf(c_0_115, negated_conjecture, (r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_110]), c_0_79])).
cnf(c_0_116, negated_conjecture, (k4_xboole_0(esk3_0,esk1_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_111]), c_0_77]), c_0_25]), c_0_103])).
cnf(c_0_117, negated_conjecture, (k4_xboole_0(esk1_0,esk3_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_112, c_0_83])).
cnf(c_0_118, negated_conjecture, (~r2_xboole_0(esk2_0,X1)|~r2_hidden(esk7_2(esk2_0,X1),esk4_0)), inference(spm,[status(thm)],[c_0_113, c_0_114])).
cnf(c_0_119, plain, (r2_hidden(esk7_2(X1,X2),X2)|~r2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_107])).
fof(c_0_120, plain, ![X57, X58]:(((r1_tarski(X57,X58)|~r2_xboole_0(X57,X58))&(X57!=X58|~r2_xboole_0(X57,X58)))&(~r1_tarski(X57,X58)|X57=X58|r2_xboole_0(X57,X58))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_xboole_0])])])])).
cnf(c_0_121, negated_conjecture, (r1_tarski(esk2_0,X1)|r2_hidden(esk8_2(esk2_0,X1),esk4_0)), inference(spm,[status(thm)],[c_0_115, c_0_22])).
cnf(c_0_122, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_116]), c_0_79])).
cnf(c_0_123, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_117]), c_0_79])).
cnf(c_0_124, negated_conjecture, (~r2_xboole_0(esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_118, c_0_119])).
cnf(c_0_125, plain, (X1=X2|r2_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_126, negated_conjecture, (r1_tarski(esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_27, c_0_121])).
cnf(c_0_127, negated_conjecture, (~r2_xboole_0(esk1_0,X1)|~r2_hidden(esk7_2(esk1_0,X1),esk3_0)), inference(spm,[status(thm)],[c_0_113, c_0_122])).
cnf(c_0_128, negated_conjecture, (r1_tarski(esk1_0,X1)|r2_hidden(esk8_2(esk1_0,X1),esk3_0)), inference(spm,[status(thm)],[c_0_123, c_0_22])).
cnf(c_0_129, negated_conjecture, (esk1_0!=esk3_0|esk2_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_130, negated_conjecture, (esk2_0=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125]), c_0_126])])).
cnf(c_0_131, negated_conjecture, (~r2_xboole_0(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_127, c_0_119])).
cnf(c_0_132, negated_conjecture, (r1_tarski(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_27, c_0_128])).
cnf(c_0_133, negated_conjecture, (esk1_0!=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_129, c_0_130])])).
cnf(c_0_134, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_125]), c_0_132])]), c_0_133]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 135
# Proof object clause steps            : 104
# Proof object formula steps           : 31
# Proof object conjectures             : 57
# Proof object clause conjectures      : 54
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 72
# Proof object simplifying inferences  : 59
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 37
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 37
# Processed clauses                    : 2532
# ...of these trivial                  : 286
# ...subsumed                          : 1650
# ...remaining for further processing  : 596
# Other redundant clauses eliminated   : 9
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 13
# Backward-rewritten                   : 220
# Generated clauses                    : 26884
# ...of the previous two non-trivial   : 17123
# Contextual simplify-reflections      : 0
# Paramodulations                      : 26865
# Factorizations                       : 10
# Equation resolutions                 : 9
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 354
#    Positive orientable unit clauses  : 135
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 19
#    Non-unit-clauses                  : 197
# Current number of unprocessed clauses: 14395
# ...number of literals in the above   : 31801
# Current number of archived formulas  : 0
# Current number of archived clauses   : 233
# Clause-clause subsumption calls (NU) : 10763
# Rec. Clause-clause subsumption calls : 9080
# Non-unit clause-clause subsumptions  : 627
# Unit Clause-clause subsumption calls : 2781
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 240
# BW rewrite match successes           : 75
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 259848

# -------------------------------------------------
# User time                : 0.139 s
# System time              : 0.007 s
# Total time               : 0.146 s
# Maximum resident set size: 28500 pages
