# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t23_relat_1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>(X3=k1_tarski(k4_tarski(X1,X2))=>(k1_relat_1(X3)=k1_tarski(X1)&k2_relat_1(X3)=k1_tarski(X2)))), file('data/MPTP/proofs/2bd7d9da97b54d31bc89432d931c9188.E_input', t23_relat_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/2bd7d9da97b54d31bc89432d931c9188.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/2bd7d9da97b54d31bc89432d931c9188.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/2bd7d9da97b54d31bc89432d931c9188.E_input', t71_enumset1)).
fof(t35_zfmisc_1, axiom, ![X1, X2]:k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2))=k1_tarski(k4_tarski(X1,X2)), file('data/MPTP/proofs/2bd7d9da97b54d31bc89432d931c9188.E_input', t35_zfmisc_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/2bd7d9da97b54d31bc89432d931c9188.E_input', d1_tarski)).
fof(t129_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,k1_tarski(X4)))<=>(r2_hidden(X1,X3)&X2=X4)), file('data/MPTP/proofs/2bd7d9da97b54d31bc89432d931c9188.E_input', t129_zfmisc_1)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/2bd7d9da97b54d31bc89432d931c9188.E_input', d4_relat_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/2bd7d9da97b54d31bc89432d931c9188.E_input', d5_relat_1)).
fof(c_0_9, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>(X3=k1_tarski(k4_tarski(X1,X2))=>(k1_relat_1(X3)=k1_tarski(X1)&k2_relat_1(X3)=k1_tarski(X2))))), inference(assume_negation,[status(cth)],[t23_relat_1])).
fof(c_0_10, negated_conjecture, (v1_relat_1(esk3_0)&(esk3_0=k1_tarski(k4_tarski(esk1_0,esk2_0))&(k1_relat_1(esk3_0)!=k1_tarski(esk1_0)|k2_relat_1(esk3_0)!=k1_tarski(esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
fof(c_0_11, plain, ![X8]:k2_tarski(X8,X8)=k1_tarski(X8), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_12, plain, ![X9, X10]:k1_enumset1(X9,X9,X10)=k2_tarski(X9,X10), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_13, plain, ![X11, X12, X13]:k2_enumset1(X11,X11,X12,X13)=k1_enumset1(X11,X12,X13), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_14, plain, ![X14, X15]:k2_zfmisc_1(k1_tarski(X14),k1_tarski(X15))=k1_tarski(k4_tarski(X14,X15)), inference(variable_rename,[status(thm)],[t35_zfmisc_1])).
fof(c_0_15, plain, ![X16, X17, X18, X19, X20, X21]:(((~r2_hidden(X18,X17)|X18=X16|X17!=k1_tarski(X16))&(X19!=X16|r2_hidden(X19,X17)|X17!=k1_tarski(X16)))&((~r2_hidden(esk4_2(X20,X21),X21)|esk4_2(X20,X21)!=X20|X21=k1_tarski(X20))&(r2_hidden(esk4_2(X20,X21),X21)|esk4_2(X20,X21)=X20|X21=k1_tarski(X20)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_16, plain, ![X23, X24, X25, X26]:(((r2_hidden(X23,X25)|~r2_hidden(k4_tarski(X23,X24),k2_zfmisc_1(X25,k1_tarski(X26))))&(X24=X26|~r2_hidden(k4_tarski(X23,X24),k2_zfmisc_1(X25,k1_tarski(X26)))))&(~r2_hidden(X23,X25)|X24!=X26|r2_hidden(k4_tarski(X23,X24),k2_zfmisc_1(X25,k1_tarski(X26))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t129_zfmisc_1])])])).
cnf(c_0_17, negated_conjecture, (esk3_0=k1_tarski(k4_tarski(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2))=k1_tarski(k4_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (X1=X2|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,k1_tarski(X2)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, negated_conjecture, (esk3_0=k2_enumset1(k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_18]), c_0_19]), c_0_20])).
cnf(c_0_25, plain, (k2_zfmisc_1(k2_enumset1(X1,X1,X1,X1),k2_enumset1(X2,X2,X2,X2))=k2_enumset1(k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_18]), c_0_18]), c_0_18]), c_0_19]), c_0_19]), c_0_19]), c_0_20]), c_0_20]), c_0_20])).
cnf(c_0_26, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_enumset1(X2,X2,X2,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_18]), c_0_19]), c_0_20])).
cnf(c_0_27, plain, (X1=X2|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,k2_enumset1(X2,X2,X2,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_18]), c_0_19]), c_0_20])).
cnf(c_0_28, negated_conjecture, (k2_zfmisc_1(k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))=esk3_0), inference(rw,[status(thm)],[c_0_24, c_0_25])).
fof(c_0_29, plain, ![X27, X28, X29, X31, X32, X33, X34, X36]:(((~r2_hidden(X29,X28)|r2_hidden(k4_tarski(X29,esk5_3(X27,X28,X29)),X27)|X28!=k1_relat_1(X27))&(~r2_hidden(k4_tarski(X31,X32),X27)|r2_hidden(X31,X28)|X28!=k1_relat_1(X27)))&((~r2_hidden(esk6_2(X33,X34),X34)|~r2_hidden(k4_tarski(esk6_2(X33,X34),X36),X33)|X34=k1_relat_1(X33))&(r2_hidden(esk6_2(X33,X34),X34)|r2_hidden(k4_tarski(esk6_2(X33,X34),esk7_2(X33,X34)),X33)|X34=k1_relat_1(X33)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
cnf(c_0_30, plain, (r2_hidden(X1,X2)|X2!=k2_enumset1(X1,X1,X1,X1)), inference(er,[status(thm)],[c_0_26])).
fof(c_0_31, plain, ![X38, X39, X40, X42, X43, X44, X45, X47]:(((~r2_hidden(X40,X39)|r2_hidden(k4_tarski(esk8_3(X38,X39,X40),X40),X38)|X39!=k2_relat_1(X38))&(~r2_hidden(k4_tarski(X43,X42),X38)|r2_hidden(X42,X39)|X39!=k2_relat_1(X38)))&((~r2_hidden(esk9_2(X44,X45),X45)|~r2_hidden(k4_tarski(X47,esk9_2(X44,X45)),X44)|X45=k2_relat_1(X44))&(r2_hidden(esk9_2(X44,X45),X45)|r2_hidden(k4_tarski(esk10_2(X44,X45),esk9_2(X44,X45)),X44)|X45=k2_relat_1(X44)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_32, negated_conjecture, (X1=esk2_0|~r2_hidden(k4_tarski(X2,X1),esk3_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_33, plain, (r2_hidden(k4_tarski(X1,esk5_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34, plain, (r2_hidden(k4_tarski(X1,X2),X3)|X3!=k2_zfmisc_1(k2_enumset1(X1,X1,X1,X1),k2_enumset1(X2,X2,X2,X2))), inference(spm,[status(thm)],[c_0_30, c_0_25])).
cnf(c_0_35, plain, (r2_hidden(k4_tarski(esk8_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_36, negated_conjecture, (esk5_3(esk3_0,X1,X2)=esk2_0|X1!=k1_relat_1(esk3_0)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_37, plain, (r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),X1)|X1!=esk3_0), inference(spm,[status(thm)],[c_0_34, c_0_28])).
cnf(c_0_39, negated_conjecture, (X1=esk2_0|X2!=k2_relat_1(esk3_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_32, c_0_35])).
cnf(c_0_40, plain, (r2_hidden(esk4_2(X1,X2),X2)|esk4_2(X1,X2)=X1|X2=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_41, negated_conjecture, (esk5_3(esk3_0,k1_relat_1(esk3_0),X1)=esk2_0|~r2_hidden(X1,k1_relat_1(esk3_0))), inference(er,[status(thm)],[c_0_36])).
cnf(c_0_42, negated_conjecture, (r2_hidden(esk1_0,X1)|X1!=k1_relat_1(X2)|X2!=esk3_0), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_43, negated_conjecture, (X1=esk2_0|~r2_hidden(X1,k2_relat_1(esk3_0))), inference(er,[status(thm)],[c_0_39])).
cnf(c_0_44, plain, (esk4_2(X1,X2)=X1|X2=k2_enumset1(X1,X1,X1,X1)|r2_hidden(esk4_2(X1,X2),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_18]), c_0_19]), c_0_20])).
cnf(c_0_45, negated_conjecture, (r2_hidden(k4_tarski(X1,esk2_0),esk3_0)|~r2_hidden(X1,k1_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_33, c_0_41])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(X1))|X1!=esk3_0), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_47, negated_conjecture, (k2_enumset1(X1,X1,X1,X1)=k2_relat_1(esk3_0)|esk4_2(X1,k2_relat_1(esk3_0))=esk2_0|esk4_2(X1,k2_relat_1(esk3_0))=X1), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_48, plain, (r2_hidden(X2,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_49, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_50, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,k1_tarski(X4)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_51, plain, (X2=k1_tarski(X1)|~r2_hidden(esk4_2(X1,X2),X2)|esk4_2(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_52, negated_conjecture, (k2_enumset1(X1,X1,X1,X1)=k2_relat_1(esk3_0)|esk4_2(X1,k2_relat_1(esk3_0))=esk2_0|X1!=esk2_0), inference(ef,[status(thm)],[c_0_47])).
cnf(c_0_53, negated_conjecture, (r2_hidden(esk2_0,X1)|X1!=k2_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_54, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,k2_enumset1(X4,X4,X4,X4)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_18]), c_0_19]), c_0_20])).
cnf(c_0_55, negated_conjecture, (k1_relat_1(esk3_0)!=k1_tarski(esk1_0)|k2_relat_1(esk3_0)!=k1_tarski(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_56, plain, (X2=k2_enumset1(X1,X1,X1,X1)|esk4_2(X1,X2)!=X1|~r2_hidden(esk4_2(X1,X2),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_18]), c_0_19]), c_0_20])).
cnf(c_0_57, negated_conjecture, (k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)=k2_relat_1(esk3_0)|esk4_2(esk2_0,k2_relat_1(esk3_0))=esk2_0), inference(er,[status(thm)],[c_0_52])).
cnf(c_0_58, negated_conjecture, (r2_hidden(esk2_0,k2_relat_1(esk3_0))), inference(er,[status(thm)],[c_0_53])).
cnf(c_0_59, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_60, negated_conjecture, (r2_hidden(X1,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))|~r2_hidden(k4_tarski(X1,X2),esk3_0)), inference(spm,[status(thm)],[c_0_54, c_0_28])).
cnf(c_0_61, negated_conjecture, (k1_relat_1(esk3_0)!=k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)|k2_relat_1(esk3_0)!=k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_18]), c_0_18]), c_0_19]), c_0_19]), c_0_20]), c_0_20])).
cnf(c_0_62, negated_conjecture, (k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)=k2_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58])])).
cnf(c_0_63, plain, (X1=X3|X2!=k2_enumset1(X3,X3,X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_18]), c_0_19]), c_0_20])).
cnf(c_0_64, negated_conjecture, (r2_hidden(X1,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))|X2!=k1_relat_1(esk3_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_60, c_0_33])).
cnf(c_0_65, negated_conjecture, (k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)!=k1_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_62])])).
cnf(c_0_66, plain, (X1=X2|~r2_hidden(X1,k2_enumset1(X2,X2,X2,X2))), inference(er,[status(thm)],[c_0_63])).
cnf(c_0_67, negated_conjecture, (r2_hidden(X1,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))|~r2_hidden(X1,k1_relat_1(esk3_0))), inference(er,[status(thm)],[c_0_64])).
cnf(c_0_68, negated_conjecture, (esk4_2(esk1_0,X1)=esk1_0|r2_hidden(esk4_2(esk1_0,X1),X1)|X1!=k1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_65, c_0_44])).
cnf(c_0_69, negated_conjecture, (X1=esk1_0|~r2_hidden(X1,k1_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_70, negated_conjecture, (r2_hidden(esk1_0,X1)|X1!=k1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_37, c_0_49])).
cnf(c_0_71, negated_conjecture, (esk4_2(esk1_0,k1_relat_1(esk3_0))=esk1_0), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_68]), c_0_69])).
cnf(c_0_72, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(esk3_0))), inference(er,[status(thm)],[c_0_70])).
cnf(c_0_73, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_71]), c_0_72])]), c_0_65]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 74
# Proof object clause steps            : 55
# Proof object formula steps           : 19
# Proof object conjectures             : 34
# Proof object clause conjectures      : 31
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 46
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 19
# Processed clauses                    : 930
# ...of these trivial                  : 34
# ...subsumed                          : 377
# ...remaining for further processing  : 519
# Other redundant clauses eliminated   : 259
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 82
# Backward-rewritten                   : 23
# Generated clauses                    : 14146
# ...of the previous two non-trivial   : 13421
# Contextual simplify-reflections      : 11
# Paramodulations                      : 13778
# Factorizations                       : 18
# Equation resolutions                 : 341
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 384
#    Positive orientable unit clauses  : 39
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 344
# Current number of unprocessed clauses: 12468
# ...number of literals in the above   : 67513
# Current number of archived formulas  : 0
# Current number of archived clauses   : 136
# Clause-clause subsumption calls (NU) : 13562
# Rec. Clause-clause subsumption calls : 7394
# Non-unit clause-clause subsumptions  : 459
# Unit Clause-clause subsumption calls : 1920
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 179
# BW rewrite match successes           : 8
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 262596

# -------------------------------------------------
# User time                : 0.208 s
# System time              : 0.010 s
# Total time               : 0.218 s
# Maximum resident set size: 28764 pages
