# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t14_tex_2, conjecture, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X1)=>((g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))=g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))&v1_tex_2(X2,X1))=>v1_tex_2(X3,X1))))), file('data/MPTP/proofs/2c2daac7196945698c8b6214d942c33e.E_input', t14_tex_2)).
fof(d3_tex_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>(v1_tex_2(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>v1_subset_1(X3,u1_struct_0(X1))))))), file('data/MPTP/proofs/2c2daac7196945698c8b6214d942c33e.E_input', d3_tex_2)).
fof(d7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(v1_subset_1(X2,X1)<=>~(X2=X1))), file('data/MPTP/proofs/2c2daac7196945698c8b6214d942c33e.E_input', d7_subset_1)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/2c2daac7196945698c8b6214d942c33e.E_input', dt_m1_pre_topc)).
fof(t11_tmap_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>(v1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))&m1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)),X1)))), file('data/MPTP/proofs/2c2daac7196945698c8b6214d942c33e.E_input', t11_tmap_1)).
fof(t2_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>m1_pre_topc(X1,X1)), file('data/MPTP/proofs/2c2daac7196945698c8b6214d942c33e.E_input', t2_tsep_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/2c2daac7196945698c8b6214d942c33e.E_input', d10_xboole_0)).
fof(t59_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))=>m1_pre_topc(X2,X1))), file('data/MPTP/proofs/2c2daac7196945698c8b6214d942c33e.E_input', t59_pre_topc)).
fof(fc14_subset_1, axiom, ![X1]:~(v1_subset_1(k2_subset_1(X1),X1)), file('data/MPTP/proofs/2c2daac7196945698c8b6214d942c33e.E_input', fc14_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/2c2daac7196945698c8b6214d942c33e.E_input', d4_subset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/2c2daac7196945698c8b6214d942c33e.E_input', t3_subset)).
fof(t5_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X1)=>(u1_struct_0(X2)=u1_struct_0(X3)=>g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))=g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3)))))), file('data/MPTP/proofs/2c2daac7196945698c8b6214d942c33e.E_input', t5_tsep_1)).
fof(dt_k1_pre_topc, axiom, ![X1, X2]:((l1_pre_topc(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_pre_topc(k1_pre_topc(X1,X2))&m1_pre_topc(k1_pre_topc(X1,X2),X1))), file('data/MPTP/proofs/2c2daac7196945698c8b6214d942c33e.E_input', dt_k1_pre_topc)).
fof(t29_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>u1_struct_0(k1_pre_topc(X1,X2))=X2)), file('data/MPTP/proofs/2c2daac7196945698c8b6214d942c33e.E_input', t29_pre_topc)).
fof(abstractness_v1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v1_pre_topc(X1)=>X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))), file('data/MPTP/proofs/2c2daac7196945698c8b6214d942c33e.E_input', abstractness_v1_pre_topc)).
fof(t35_borsuk_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>r1_tarski(u1_struct_0(X2),u1_struct_0(X1)))), file('data/MPTP/proofs/2c2daac7196945698c8b6214d942c33e.E_input', t35_borsuk_1)).
fof(c_0_16, negated_conjecture, ~(![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X1)=>((g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))=g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))&v1_tex_2(X2,X1))=>v1_tex_2(X3,X1)))))), inference(assume_negation,[status(cth)],[t14_tex_2])).
fof(c_0_17, plain, ![X7, X8, X9]:((~v1_tex_2(X8,X7)|(~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X7)))|(X9!=u1_struct_0(X8)|v1_subset_1(X9,u1_struct_0(X7))))|~m1_pre_topc(X8,X7)|~l1_pre_topc(X7))&((m1_subset_1(esk4_2(X7,X8),k1_zfmisc_1(u1_struct_0(X7)))|v1_tex_2(X8,X7)|~m1_pre_topc(X8,X7)|~l1_pre_topc(X7))&((esk4_2(X7,X8)=u1_struct_0(X8)|v1_tex_2(X8,X7)|~m1_pre_topc(X8,X7)|~l1_pre_topc(X7))&(~v1_subset_1(esk4_2(X7,X8),u1_struct_0(X7))|v1_tex_2(X8,X7)|~m1_pre_topc(X8,X7)|~l1_pre_topc(X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tex_2])])])])])).
fof(c_0_18, plain, ![X20, X21]:((~v1_subset_1(X21,X20)|X21!=X20|~m1_subset_1(X21,k1_zfmisc_1(X20)))&(X21=X20|v1_subset_1(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(X20)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_subset_1])])])])).
fof(c_0_19, plain, ![X28, X29]:(~l1_pre_topc(X28)|(~m1_pre_topc(X29,X28)|l1_pre_topc(X29))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
fof(c_0_20, plain, ![X24, X25]:((v1_pre_topc(g1_pre_topc(u1_struct_0(X25),u1_pre_topc(X25)))|~m1_pre_topc(X25,X24)|~l1_pre_topc(X24))&(m1_pre_topc(g1_pre_topc(u1_struct_0(X25),u1_pre_topc(X25)),X24)|~m1_pre_topc(X25,X24)|~l1_pre_topc(X24))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_tmap_1])])])])).
fof(c_0_21, plain, ![X30]:(~l1_pre_topc(X30)|m1_pre_topc(X30,X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tsep_1])])).
fof(c_0_22, negated_conjecture, (l1_pre_topc(esk1_0)&(m1_pre_topc(esk2_0,esk1_0)&(m1_pre_topc(esk3_0,esk1_0)&((g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))=g1_pre_topc(u1_struct_0(esk3_0),u1_pre_topc(esk3_0))&v1_tex_2(esk2_0,esk1_0))&~v1_tex_2(esk3_0,esk1_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).
cnf(c_0_23, plain, (v1_tex_2(X2,X1)|~v1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (X1=X2|v1_subset_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|v1_tex_2(X2,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (m1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (m1_pre_topc(X1,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, plain, (esk4_2(X1,X2)=u1_struct_0(X1)|v1_tex_2(X2,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
cnf(c_0_31, plain, (esk4_2(X1,X2)=u1_struct_0(X2)|v1_tex_2(X2,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_32, plain, (l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_33, negated_conjecture, (m1_pre_topc(esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34, negated_conjecture, (~v1_tex_2(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_35, negated_conjecture, (esk4_2(esk1_0,X1)=u1_struct_0(esk1_0)|v1_tex_2(X1,esk1_0)|~m1_pre_topc(X1,esk1_0)), inference(spm,[status(thm)],[c_0_30, c_0_29])).
cnf(c_0_36, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_37, plain, ![X12, X13]:(((r1_tarski(X12,X13)|X12!=X13)&(r1_tarski(X13,X12)|X12!=X13))&(~r1_tarski(X12,X13)|~r1_tarski(X13,X12)|X12=X13)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_38, negated_conjecture, (u1_struct_0(X1)=esk4_2(esk1_0,X1)|v1_tex_2(X1,esk1_0)|~m1_pre_topc(X1,esk1_0)), inference(spm,[status(thm)],[c_0_31, c_0_29])).
fof(c_0_39, plain, ![X31, X32]:(~l1_pre_topc(X31)|(~m1_pre_topc(X32,g1_pre_topc(u1_struct_0(X31),u1_pre_topc(X31)))|m1_pre_topc(X32,X31))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t59_pre_topc])])])).
cnf(c_0_40, negated_conjecture, (l1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_29])])).
fof(c_0_41, plain, ![X33]:~v1_subset_1(k2_subset_1(X33),X33), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc14_subset_1])])).
fof(c_0_42, plain, ![X19]:k2_subset_1(X19)=X19, inference(variable_rename,[status(thm)],[d4_subset_1])).
cnf(c_0_43, plain, (v1_subset_1(X3,u1_struct_0(X2))|~v1_tex_2(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|X3!=u1_struct_0(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_44, negated_conjecture, (u1_struct_0(esk1_0)=esk4_2(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])])).
fof(c_0_45, plain, ![X34, X35]:((~m1_subset_1(X34,k1_zfmisc_1(X35))|r1_tarski(X34,X35))&(~r1_tarski(X34,X35)|m1_subset_1(X34,k1_zfmisc_1(X35)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_46, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_47, plain, ![X14, X15, X16]:(~l1_pre_topc(X14)|(~m1_pre_topc(X15,X14)|(~m1_pre_topc(X16,X14)|(u1_struct_0(X15)!=u1_struct_0(X16)|g1_pre_topc(u1_struct_0(X15),u1_pre_topc(X15))=g1_pre_topc(u1_struct_0(X16),u1_pre_topc(X16)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_tsep_1])])])).
fof(c_0_48, plain, ![X26, X27]:((v1_pre_topc(k1_pre_topc(X26,X27))|(~l1_pre_topc(X26)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X26)))))&(m1_pre_topc(k1_pre_topc(X26,X27),X26)|(~l1_pre_topc(X26)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X26)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_pre_topc])])])).
cnf(c_0_49, negated_conjecture, (g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))=g1_pre_topc(u1_struct_0(esk3_0),u1_pre_topc(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_50, negated_conjecture, (u1_struct_0(esk3_0)=esk4_2(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_38]), c_0_36])])).
cnf(c_0_51, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_52, plain, (m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_53, negated_conjecture, (m1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))), inference(spm,[status(thm)],[c_0_28, c_0_40])).
cnf(c_0_54, plain, (~v1_subset_1(k2_subset_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_55, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_56, negated_conjecture, (u1_struct_0(X1)=esk4_2(esk1_0,X1)|v1_subset_1(X2,esk4_2(esk1_0,esk3_0))|X2!=u1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(esk4_2(esk1_0,esk3_0)))|~m1_pre_topc(X1,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_38]), c_0_29])]), c_0_44]), c_0_44])).
cnf(c_0_57, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_58, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_46])).
cnf(c_0_59, plain, (g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))=g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|u1_struct_0(X2)!=u1_struct_0(X3)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_60, plain, (m1_pre_topc(k1_pre_topc(X1,X2),X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_61, plain, ![X22, X23]:(~l1_pre_topc(X22)|(~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(X22)))|u1_struct_0(k1_pre_topc(X22,X23))=X23)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t29_pre_topc])])])).
cnf(c_0_62, negated_conjecture, (g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))=g1_pre_topc(esk4_2(esk1_0,esk3_0),u1_pre_topc(esk3_0))), inference(rw,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_63, negated_conjecture, (l1_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_51]), c_0_29])])).
cnf(c_0_64, negated_conjecture, (l1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_36]), c_0_49]), c_0_29])])).
cnf(c_0_65, negated_conjecture, (m1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_29])])).
cnf(c_0_66, plain, (~v1_subset_1(X1,X1)), inference(rw,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_67, negated_conjecture, (esk4_2(esk1_0,esk3_0)=esk4_2(esk1_0,esk1_0)|v1_subset_1(X1,esk4_2(esk1_0,esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(esk4_2(esk1_0,esk3_0)))), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_33]), c_0_44]), c_0_44]), c_0_24])).
cnf(c_0_68, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_69, plain, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=g1_pre_topc(u1_struct_0(k1_pre_topc(X2,X3)),u1_pre_topc(k1_pre_topc(X2,X3)))|u1_struct_0(X1)!=u1_struct_0(k1_pre_topc(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_70, plain, (u1_struct_0(k1_pre_topc(X1,X2))=X2|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_71, negated_conjecture, (l1_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_36]), c_0_29])])).
cnf(c_0_72, negated_conjecture, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=g1_pre_topc(esk4_2(esk1_0,esk3_0),u1_pre_topc(esk3_0))|u1_struct_0(X1)!=esk4_2(esk1_0,esk3_0)|~m1_pre_topc(X1,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_36]), c_0_49]), c_0_29])]), c_0_62]), c_0_50])).
cnf(c_0_73, negated_conjecture, (u1_struct_0(X1)=esk4_2(esk2_0,X1)|v1_tex_2(X1,esk2_0)|~m1_pre_topc(X1,esk2_0)), inference(spm,[status(thm)],[c_0_31, c_0_63])).
cnf(c_0_74, negated_conjecture, (l1_pre_topc(g1_pre_topc(esk4_2(esk1_0,esk3_0),u1_pre_topc(esk3_0)))), inference(rw,[status(thm)],[c_0_64, c_0_62])).
cnf(c_0_75, negated_conjecture, (m1_pre_topc(g1_pre_topc(esk4_2(esk1_0,esk3_0),u1_pre_topc(esk1_0)),esk1_0)), inference(rw,[status(thm)],[c_0_65, c_0_44])).
cnf(c_0_76, negated_conjecture, (esk4_2(esk1_0,esk3_0)=esk4_2(esk1_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68])])).
cnf(c_0_77, plain, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X3,X2)))|u1_struct_0(X1)!=X2|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_pre_topc(X1,X3)|~l1_pre_topc(X3)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_78, negated_conjecture, (m1_pre_topc(esk3_0,esk3_0)), inference(spm,[status(thm)],[c_0_28, c_0_71])).
cnf(c_0_79, negated_conjecture, (g1_pre_topc(esk4_2(esk1_0,esk3_0),u1_pre_topc(esk3_0))=g1_pre_topc(esk4_2(esk1_0,esk3_0),u1_pre_topc(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_44]), c_0_33])])).
fof(c_0_80, plain, ![X11]:(~l1_pre_topc(X11)|(~v1_pre_topc(X11)|X11=g1_pre_topc(u1_struct_0(X11),u1_pre_topc(X11)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_pre_topc])])).
cnf(c_0_81, negated_conjecture, (u1_struct_0(X1)=esk4_2(esk2_0,X1)|v1_subset_1(X2,u1_struct_0(esk2_0))|X2!=u1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_73]), c_0_63])])).
cnf(c_0_82, negated_conjecture, (m1_pre_topc(esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_28, c_0_63])).
cnf(c_0_83, negated_conjecture, (m1_pre_topc(X1,esk2_0)|~m1_pre_topc(X1,g1_pre_topc(esk4_2(esk1_0,esk3_0),u1_pre_topc(esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_62]), c_0_63])])).
cnf(c_0_84, negated_conjecture, (m1_pre_topc(g1_pre_topc(esk4_2(esk1_0,esk3_0),u1_pre_topc(esk3_0)),g1_pre_topc(esk4_2(esk1_0,esk3_0),u1_pre_topc(esk3_0)))), inference(spm,[status(thm)],[c_0_28, c_0_74])).
cnf(c_0_85, negated_conjecture, (m1_pre_topc(g1_pre_topc(esk4_2(esk1_0,esk1_0),u1_pre_topc(esk1_0)),esk1_0)), inference(rw,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_86, negated_conjecture, (g1_pre_topc(esk4_2(esk1_0,esk1_0),u1_pre_topc(esk1_0))=g1_pre_topc(X1,u1_pre_topc(k1_pre_topc(esk3_0,X1)))|esk4_2(esk1_0,esk1_0)!=X1|~m1_subset_1(X1,k1_zfmisc_1(esk4_2(esk1_0,esk1_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_50]), c_0_79]), c_0_50]), c_0_50]), c_0_71])]), c_0_76]), c_0_76]), c_0_76])).
cnf(c_0_87, plain, (X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X1)|~v1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_88, plain, (v1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_89, plain, ![X17, X18]:(~l1_pre_topc(X17)|(~m1_pre_topc(X18,X17)|r1_tarski(u1_struct_0(X18),u1_struct_0(X17)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t35_borsuk_1])])])).
cnf(c_0_90, negated_conjecture, (v1_tex_2(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_91, negated_conjecture, (u1_struct_0(esk2_0)=esk4_2(esk2_0,esk2_0)|v1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_24])).
cnf(c_0_92, negated_conjecture, (m1_pre_topc(g1_pre_topc(esk4_2(esk1_0,esk3_0),u1_pre_topc(esk3_0)),esk2_0)), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_93, negated_conjecture, (m1_pre_topc(g1_pre_topc(X1,u1_pre_topc(k1_pre_topc(esk3_0,X1))),esk1_0)|esk4_2(esk1_0,esk1_0)!=X1|~m1_subset_1(X1,k1_zfmisc_1(esk4_2(esk1_0,esk1_0)))), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_94, plain, (g1_pre_topc(X1,u1_pre_topc(k1_pre_topc(X2,X1)))=k1_pre_topc(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(k1_pre_topc(X2,X1))|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_70]), c_0_88])).
cnf(c_0_95, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_96, plain, (r1_tarski(u1_struct_0(X2),u1_struct_0(X1))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_97, negated_conjecture, (v1_subset_1(X1,esk4_2(esk1_0,esk3_0))|X1!=u1_struct_0(esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(esk4_2(esk1_0,esk3_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_90]), c_0_51]), c_0_29])]), c_0_44]), c_0_44])).
cnf(c_0_98, negated_conjecture, (u1_struct_0(esk2_0)=esk4_2(esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_91]), c_0_68])])).
cnf(c_0_99, negated_conjecture, (m1_pre_topc(g1_pre_topc(esk4_2(esk1_0,esk3_0),u1_pre_topc(esk1_0)),esk2_0)), inference(rw,[status(thm)],[c_0_92, c_0_79])).
cnf(c_0_100, negated_conjecture, (m1_pre_topc(k1_pre_topc(esk3_0,X1),esk1_0)|esk4_2(esk1_0,esk1_0)!=X1|~m1_subset_1(X1,k1_zfmisc_1(esk4_2(esk1_0,esk1_0)))|~l1_pre_topc(k1_pre_topc(esk3_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_50]), c_0_76]), c_0_71])])).
cnf(c_0_101, plain, (l1_pre_topc(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_26, c_0_60])).
cnf(c_0_102, plain, (u1_struct_0(X1)=u1_struct_0(X2)|~r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_103, negated_conjecture, (v1_subset_1(X1,esk4_2(esk1_0,esk3_0))|X1!=esk4_2(esk2_0,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(esk4_2(esk1_0,esk3_0)))), inference(rw,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_104, negated_conjecture, (m1_pre_topc(g1_pre_topc(esk4_2(esk1_0,esk1_0),u1_pre_topc(esk1_0)),esk2_0)), inference(rw,[status(thm)],[c_0_99, c_0_76])).
cnf(c_0_105, negated_conjecture, (m1_pre_topc(k1_pre_topc(esk3_0,X1),esk1_0)|esk4_2(esk1_0,esk1_0)!=X1|~m1_subset_1(X1,k1_zfmisc_1(esk4_2(esk1_0,esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_50]), c_0_76]), c_0_71])])).
cnf(c_0_106, negated_conjecture, (u1_struct_0(X1)=esk4_2(esk2_0,esk2_0)|~r1_tarski(u1_struct_0(X1),esk4_2(esk2_0,esk2_0))|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_102, c_0_98])).
cnf(c_0_107, negated_conjecture, (esk4_2(esk2_0,esk2_0)!=esk4_2(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_103]), c_0_68])])).
cnf(c_0_108, negated_conjecture, (m1_pre_topc(g1_pre_topc(X1,u1_pre_topc(k1_pre_topc(esk3_0,X1))),esk2_0)|esk4_2(esk1_0,esk1_0)!=X1|~m1_subset_1(X1,k1_zfmisc_1(esk4_2(esk1_0,esk1_0)))), inference(spm,[status(thm)],[c_0_104, c_0_86])).
cnf(c_0_109, negated_conjecture, (l1_pre_topc(k1_pre_topc(esk3_0,X1))|esk4_2(esk1_0,esk1_0)!=X1|~m1_subset_1(X1,k1_zfmisc_1(esk4_2(esk1_0,esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_105]), c_0_29])])).
cnf(c_0_110, negated_conjecture, (~r1_tarski(esk4_2(esk1_0,esk3_0),esk4_2(esk2_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_44]), c_0_51]), c_0_29])]), c_0_107])).
cnf(c_0_111, plain, (r1_tarski(X1,u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X3)))|~m1_pre_topc(k1_pre_topc(X3,X1),X2)|~l1_pre_topc(X2)|~l1_pre_topc(X3)), inference(spm,[status(thm)],[c_0_96, c_0_70])).
cnf(c_0_112, negated_conjecture, (m1_pre_topc(k1_pre_topc(esk3_0,X1),esk2_0)|esk4_2(esk1_0,esk1_0)!=X1|~m1_subset_1(X1,k1_zfmisc_1(esk4_2(esk1_0,esk1_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_94]), c_0_50]), c_0_76]), c_0_71])]), c_0_109])).
cnf(c_0_113, negated_conjecture, (~r1_tarski(esk4_2(esk1_0,esk1_0),esk4_2(esk2_0,esk2_0))), inference(rw,[status(thm)],[c_0_110, c_0_76])).
cnf(c_0_114, negated_conjecture, (r1_tarski(X1,esk4_2(esk2_0,esk2_0))|esk4_2(esk1_0,esk1_0)!=X1|~m1_subset_1(X1,k1_zfmisc_1(esk4_2(esk1_0,esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_98]), c_0_50]), c_0_76]), c_0_63]), c_0_71])])).
cnf(c_0_115, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_68])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 116
# Proof object clause steps            : 83
# Proof object formula steps           : 33
# Proof object conjectures             : 55
# Proof object clause conjectures      : 52
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 47
# Proof object simplifying inferences  : 93
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 29
# Processed clauses                    : 13940
# ...of these trivial                  : 340
# ...subsumed                          : 10844
# ...remaining for further processing  : 2756
# Other redundant clauses eliminated   : 24
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 90
# Backward-rewritten                   : 632
# Generated clauses                    : 101404
# ...of the previous two non-trivial   : 91795
# Contextual simplify-reflections      : 478
# Paramodulations                      : 101308
# Factorizations                       : 0
# Equation resolutions                 : 96
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2004
#    Positive orientable unit clauses  : 64
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 13
#    Non-unit-clauses                  : 1927
# Current number of unprocessed clauses: 70794
# ...number of literals in the above   : 422968
# Current number of archived formulas  : 0
# Current number of archived clauses   : 750
# Clause-clause subsumption calls (NU) : 582281
# Rec. Clause-clause subsumption calls : 107449
# Non-unit clause-clause subsumptions  : 7453
# Unit Clause-clause subsumption calls : 6536
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 321
# BW rewrite match successes           : 22
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5869701

# -------------------------------------------------
# User time                : 1.840 s
# System time              : 0.030 s
# Total time               : 1.870 s
# Maximum resident set size: 65080 pages
