# No SInE strategy applied
# Auto-Mode selected heuristic G_E___211_C18_F1_AE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', t5_subset)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', t4_subset_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', d3_tarski)).
fof(rc2_subset_1, axiom, ![X1]:?[X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))&v1_xboole_0(X2)), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', rc2_subset_1)).
fof(cc2_tex_2, axiom, ![X1]:((~(v1_xboole_0(X1))&v1_zfmisc_1(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(~(v1_xboole_0(X2))=>~(v1_subset_1(X2,X1))))), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', cc2_tex_2)).
fof(cc4_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>~(v1_subset_1(X2,X1)))), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', cc4_subset_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', d10_xboole_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', d1_xboole_0)).
fof(dt_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', dt_k6_domain_1)).
fof(t6_tex_2, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,X1)=>~((v1_subset_1(k6_domain_1(X1,X2),X1)&v1_zfmisc_1(X1))))), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', t6_tex_2)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', redefinition_k6_domain_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', t75_enumset1)).
fof(d6_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:(X9=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)<=>![X10]:(r2_hidden(X10,X9)<=>~((~(X10=X1)&(~(X10=X2)&(~(X10=X3)&(~(X10=X4)&(~(X10=X5)&(~(X10=X6)&(~(X10=X7)&~(X10=X8))))))))))), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', d6_enumset1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/2cf5f4bf3d584e60aa5f576470b0c5b0.E_input', t7_ordinal1)).
fof(c_0_20, plain, ![X61, X62, X63]:(~r2_hidden(X61,X62)|(~m1_subset_1(X62,k1_zfmisc_1(X63))|~v1_xboole_0(X63))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_21, plain, ![X60]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X60)), inference(variable_rename,[status(thm)],[t4_subset_1])).
cnf(c_0_22, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_23, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_24, plain, ![X68, X69, X70, X71, X72]:((~r1_tarski(X68,X69)|(~r2_hidden(X70,X68)|r2_hidden(X70,X69)))&((r2_hidden(esk5_2(X71,X72),X71)|r1_tarski(X71,X72))&(~r2_hidden(esk5_2(X71,X72),X72)|r1_tarski(X71,X72)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_25, plain, (~r2_hidden(X1,k1_xboole_0)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_26, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_27, plain, ![X13]:(m1_subset_1(esk3_1(X13),k1_zfmisc_1(X13))&v1_xboole_0(esk3_1(X13))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc2_subset_1])])).
fof(c_0_28, plain, ![X66, X67]:(v1_xboole_0(X66)|~v1_zfmisc_1(X66)|(~m1_subset_1(X67,k1_zfmisc_1(X66))|(v1_xboole_0(X67)|~v1_subset_1(X67,X66)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_tex_2])])])])).
fof(c_0_29, plain, ![X15, X16]:(~v1_xboole_0(X15)|(~m1_subset_1(X16,k1_zfmisc_1(X15))|~v1_subset_1(X16,X15))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc4_subset_1])])])])).
fof(c_0_30, plain, ![X88, X89]:(((r1_tarski(X88,X89)|X88!=X89)&(r1_tarski(X89,X88)|X88!=X89))&(~r1_tarski(X88,X89)|~r1_tarski(X89,X88)|X88=X89)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_31, plain, (r1_tarski(k1_xboole_0,X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_32, plain, (v1_xboole_0(esk3_1(X1))), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_33, plain, ![X74, X75, X76]:((~v1_xboole_0(X74)|~r2_hidden(X75,X74))&(r2_hidden(esk6_1(X76),X76)|v1_xboole_0(X76))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_34, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_zfmisc_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~v1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, plain, (~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~v1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_36, plain, ![X58, X59]:(v1_xboole_0(X58)|~m1_subset_1(X59,X58)|m1_subset_1(k6_domain_1(X58,X59),k1_zfmisc_1(X58))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_domain_1])])])).
fof(c_0_37, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,X1)=>~((v1_subset_1(k6_domain_1(X1,X2),X1)&v1_zfmisc_1(X1)))))), inference(assume_negation,[status(cth)],[t6_tex_2])).
cnf(c_0_38, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_40, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_41, plain, (v1_xboole_0(X1)|~v1_zfmisc_1(X2)|~v1_subset_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(csr,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_42, plain, (v1_xboole_0(X1)|m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_43, negated_conjecture, (~v1_xboole_0(esk1_0)&(m1_subset_1(esk2_0,esk1_0)&(v1_subset_1(k6_domain_1(esk1_0,esk2_0),esk1_0)&v1_zfmisc_1(esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_37])])])])).
cnf(c_0_44, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_45, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_40, c_0_26])).
cnf(c_0_46, plain, (v1_xboole_0(k6_domain_1(X1,X2))|v1_xboole_0(X1)|~v1_zfmisc_1(X1)|~v1_subset_1(k6_domain_1(X1,X2),X1)|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_47, negated_conjecture, (v1_subset_1(k6_domain_1(esk1_0,esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_48, negated_conjecture, (v1_zfmisc_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49, negated_conjecture, (m1_subset_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_50, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_51, plain, ![X78, X79]:(v1_xboole_0(X78)|~m1_subset_1(X79,X78)|k6_domain_1(X78,X79)=k1_tarski(X79)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_52, plain, ![X52]:k2_tarski(X52,X52)=k1_tarski(X52), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_53, plain, ![X80, X81]:k1_enumset1(X80,X80,X81)=k2_tarski(X80,X81), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_54, plain, ![X24, X25, X26]:k2_enumset1(X24,X24,X25,X26)=k1_enumset1(X24,X25,X26), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_55, plain, ![X27, X28, X29, X30]:k3_enumset1(X27,X27,X28,X29,X30)=k2_enumset1(X27,X28,X29,X30), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_56, plain, ![X53, X54, X55, X56, X57]:k4_enumset1(X53,X53,X54,X55,X56,X57)=k3_enumset1(X53,X54,X55,X56,X57), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_57, plain, ![X82, X83, X84, X85, X86, X87]:k5_enumset1(X82,X82,X83,X84,X85,X86,X87)=k4_enumset1(X82,X83,X84,X85,X86,X87), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_58, plain, ![X17, X18, X19, X20, X21, X22, X23]:k6_enumset1(X17,X17,X18,X19,X20,X21,X22,X23)=k5_enumset1(X17,X18,X19,X20,X21,X22,X23), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_59, plain, ![X31, X32, X33, X34, X35, X36, X37, X38, X39, X40, X41, X42, X43, X44, X45, X46, X47, X48, X49, X50]:(((~r2_hidden(X40,X39)|(X40=X31|(X40=X32|(X40=X33|(X40=X34|(X40=X35|(X40=X36|(X40=X37|X40=X38)))))))|X39!=k6_enumset1(X31,X32,X33,X34,X35,X36,X37,X38))&((X41!=X31|r2_hidden(X41,X39)|X39!=k6_enumset1(X31,X32,X33,X34,X35,X36,X37,X38))&((X41!=X32|r2_hidden(X41,X39)|X39!=k6_enumset1(X31,X32,X33,X34,X35,X36,X37,X38))&((X41!=X33|r2_hidden(X41,X39)|X39!=k6_enumset1(X31,X32,X33,X34,X35,X36,X37,X38))&((X41!=X34|r2_hidden(X41,X39)|X39!=k6_enumset1(X31,X32,X33,X34,X35,X36,X37,X38))&((X41!=X35|r2_hidden(X41,X39)|X39!=k6_enumset1(X31,X32,X33,X34,X35,X36,X37,X38))&((X41!=X36|r2_hidden(X41,X39)|X39!=k6_enumset1(X31,X32,X33,X34,X35,X36,X37,X38))&((X41!=X37|r2_hidden(X41,X39)|X39!=k6_enumset1(X31,X32,X33,X34,X35,X36,X37,X38))&(X41!=X38|r2_hidden(X41,X39)|X39!=k6_enumset1(X31,X32,X33,X34,X35,X36,X37,X38))))))))))&(((esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50)!=X42|~r2_hidden(esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50),X50)|X50=k6_enumset1(X42,X43,X44,X45,X46,X47,X48,X49))&((esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50)!=X43|~r2_hidden(esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50),X50)|X50=k6_enumset1(X42,X43,X44,X45,X46,X47,X48,X49))&((esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50)!=X44|~r2_hidden(esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50),X50)|X50=k6_enumset1(X42,X43,X44,X45,X46,X47,X48,X49))&((esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50)!=X45|~r2_hidden(esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50),X50)|X50=k6_enumset1(X42,X43,X44,X45,X46,X47,X48,X49))&((esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50)!=X46|~r2_hidden(esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50),X50)|X50=k6_enumset1(X42,X43,X44,X45,X46,X47,X48,X49))&((esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50)!=X47|~r2_hidden(esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50),X50)|X50=k6_enumset1(X42,X43,X44,X45,X46,X47,X48,X49))&((esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50)!=X48|~r2_hidden(esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50),X50)|X50=k6_enumset1(X42,X43,X44,X45,X46,X47,X48,X49))&(esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50)!=X49|~r2_hidden(esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50),X50)|X50=k6_enumset1(X42,X43,X44,X45,X46,X47,X48,X49)))))))))&(r2_hidden(esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50),X50)|(esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50)=X42|(esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50)=X43|(esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50)=X44|(esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50)=X45|(esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50)=X46|(esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50)=X47|(esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50)=X48|esk4_9(X42,X43,X44,X45,X46,X47,X48,X49,X50)=X49)))))))|X50=k6_enumset1(X42,X43,X44,X45,X46,X47,X48,X49)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_enumset1])])])])])])])).
cnf(c_0_60, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_61, negated_conjecture, (v1_xboole_0(k6_domain_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48]), c_0_49])]), c_0_50])).
cnf(c_0_62, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_63, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_64, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_65, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_66, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_67, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_68, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_69, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_58])).
fof(c_0_70, plain, ![X64, X65]:(~r2_hidden(X64,X65)|~r1_tarski(X65,X64)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_71, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k6_enumset1(X4,X5,X6,X7,X8,X9,X10,X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_72, negated_conjecture, (k6_domain_1(esk1_0,esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_73, plain, (k6_domain_1(X1,X2)=k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_63]), c_0_64]), c_0_65]), c_0_66]), c_0_67]), c_0_68]), c_0_69])).
cnf(c_0_74, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_75, plain, (r2_hidden(X1,X2)|X2!=k6_enumset1(X3,X4,X5,X6,X7,X8,X9,X1)), inference(er,[status(thm)],[c_0_71])).
cnf(c_0_76, negated_conjecture, (k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_49])]), c_0_50])).
cnf(c_0_77, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_74, c_0_39])).
cnf(c_0_78, negated_conjecture, (r2_hidden(esk2_0,X1)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_79, negated_conjecture, ($false), inference(spm,[status(thm)],[c_0_77, c_0_78]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 80
# Proof object clause steps            : 39
# Proof object formula steps           : 41
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 16
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 20
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 46
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 39
# Processed clauses                    : 57
# ...of these trivial                  : 0
# ...subsumed                          : 8
# ...remaining for further processing  : 49
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 8
# Generated clauses                    : 59
# ...of the previous two non-trivial   : 46
# Contextual simplify-reflections      : 1
# Paramodulations                      : 52
# Factorizations                       : 0
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 36
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 23
# Current number of unprocessed clauses: 24
# ...number of literals in the above   : 86
# Current number of archived formulas  : 0
# Current number of archived clauses   : 15
# Clause-clause subsumption calls (NU) : 58
# Rec. Clause-clause subsumption calls : 41
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 13
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2779

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.003 s
# Total time               : 0.010 s
# Maximum resident set size: 28720 pages
