# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PI_S0a
# and selection function SelectNegativeLiterals.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t13_ordinal1, axiom, ![X1, X2]:(r2_hidden(X1,k1_ordinal1(X2))<=>(r2_hidden(X1,X2)|X1=X2)), file('data/MPTP/proofs/2d9c14760dc540f1bc41cc045ac48819.E_input', t13_ordinal1)).
fof(d1_ordinal1, axiom, ![X1]:k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1)), file('data/MPTP/proofs/2d9c14760dc540f1bc41cc045ac48819.E_input', d1_ordinal1)).
fof(t34_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,k1_ordinal1(X2))<=>r1_ordinal1(X1,X2)))), file('data/MPTP/proofs/2d9c14760dc540f1bc41cc045ac48819.E_input', t34_ordinal1)).
fof(d2_ordinal1, axiom, ![X1]:(v1_ordinal1(X1)<=>![X2]:(r2_hidden(X2,X1)=>r1_tarski(X2,X1))), file('data/MPTP/proofs/2d9c14760dc540f1bc41cc045ac48819.E_input', d2_ordinal1)).
fof(t23_ordinal1, axiom, ![X1, X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)=>v3_ordinal1(X1))), file('data/MPTP/proofs/2d9c14760dc540f1bc41cc045ac48819.E_input', t23_ordinal1)).
fof(t41_ordinal1, conjecture, ![X1]:(v3_ordinal1(X1)=>(v4_ordinal1(X1)<=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X2,X1)=>r2_hidden(k1_ordinal1(X2),X1))))), file('data/MPTP/proofs/2d9c14760dc540f1bc41cc045ac48819.E_input', t41_ordinal1)).
fof(d4_tarski, axiom, ![X1, X2]:(X2=k3_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X3,X4)&r2_hidden(X4,X1)))), file('data/MPTP/proofs/2d9c14760dc540f1bc41cc045ac48819.E_input', d4_tarski)).
fof(redefinition_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)<=>r1_tarski(X1,X2))), file('data/MPTP/proofs/2d9c14760dc540f1bc41cc045ac48819.E_input', redefinition_r1_ordinal1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/2d9c14760dc540f1bc41cc045ac48819.E_input', d3_tarski)).
fof(t10_ordinal1, axiom, ![X1]:r2_hidden(X1,k1_ordinal1(X1)), file('data/MPTP/proofs/2d9c14760dc540f1bc41cc045ac48819.E_input', t10_ordinal1)).
fof(d6_ordinal1, axiom, ![X1]:(v4_ordinal1(X1)<=>X1=k3_tarski(X1)), file('data/MPTP/proofs/2d9c14760dc540f1bc41cc045ac48819.E_input', d6_ordinal1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/2d9c14760dc540f1bc41cc045ac48819.E_input', t7_ordinal1)).
fof(t33_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)<=>r1_ordinal1(k1_ordinal1(X1),X2)))), file('data/MPTP/proofs/2d9c14760dc540f1bc41cc045ac48819.E_input', t33_ordinal1)).
fof(t29_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>v3_ordinal1(k1_ordinal1(X1))), file('data/MPTP/proofs/2d9c14760dc540f1bc41cc045ac48819.E_input', t29_ordinal1)).
fof(c_0_14, plain, ![X40, X41]:((~r2_hidden(X40,k1_ordinal1(X41))|(r2_hidden(X40,X41)|X40=X41))&((~r2_hidden(X40,X41)|r2_hidden(X40,k1_ordinal1(X41)))&(X40!=X41|r2_hidden(X40,k1_ordinal1(X41))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_ordinal1])])])).
fof(c_0_15, plain, ![X31]:k1_ordinal1(X31)=k2_xboole_0(X31,k1_tarski(X31)), inference(variable_rename,[status(thm)],[d1_ordinal1])).
fof(c_0_16, plain, ![X42, X43]:((~r2_hidden(X42,k1_ordinal1(X43))|r1_ordinal1(X42,X43)|~v3_ordinal1(X43)|~v3_ordinal1(X42))&(~r1_ordinal1(X42,X43)|r2_hidden(X42,k1_ordinal1(X43))|~v3_ordinal1(X43)|~v3_ordinal1(X42))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t34_ordinal1])])])])).
cnf(c_0_17, plain, (r2_hidden(X1,k1_ordinal1(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_19, plain, ![X19, X20, X21]:((~v1_ordinal1(X19)|(~r2_hidden(X20,X19)|r1_tarski(X20,X19)))&((r2_hidden(esk6_1(X21),X21)|v1_ordinal1(X21))&(~r1_tarski(esk6_1(X21),X21)|v1_ordinal1(X21)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_ordinal1])])])])])])).
fof(c_0_20, plain, ![X29, X30]:(~v3_ordinal1(X30)|(~r2_hidden(X29,X30)|v3_ordinal1(X29))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_ordinal1])])).
fof(c_0_21, negated_conjecture, ~(![X1]:(v3_ordinal1(X1)=>(v4_ordinal1(X1)<=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X2,X1)=>r2_hidden(k1_ordinal1(X2),X1)))))), inference(assume_negation,[status(cth)],[t41_ordinal1])).
cnf(c_0_22, plain, (r1_ordinal1(X1,X2)|~r2_hidden(X1,k1_ordinal1(X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_24, plain, (r2_hidden(esk6_1(X1),X1)|v1_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_26, negated_conjecture, ![X7]:(v3_ordinal1(esk1_0)&(((v3_ordinal1(esk2_0)|~v4_ordinal1(esk1_0))&((r2_hidden(esk2_0,esk1_0)|~v4_ordinal1(esk1_0))&(~r2_hidden(k1_ordinal1(esk2_0),esk1_0)|~v4_ordinal1(esk1_0))))&(v4_ordinal1(esk1_0)|(~v3_ordinal1(X7)|(~r2_hidden(X7,esk1_0)|r2_hidden(k1_ordinal1(X7),esk1_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])])])).
cnf(c_0_27, plain, (r1_ordinal1(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))), inference(rw,[status(thm)],[c_0_22, c_0_18])).
cnf(c_0_28, plain, (v1_ordinal1(X1)|r2_hidden(esk6_1(X1),k2_xboole_0(X1,k1_tarski(X1)))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_29, plain, (v1_ordinal1(X1)|v3_ordinal1(esk6_1(X1))|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_25, c_0_24])).
cnf(c_0_30, negated_conjecture, (v3_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_31, plain, ![X8, X9, X10, X12, X13, X14, X15, X17]:((((r2_hidden(X10,esk3_3(X8,X9,X10))|~r2_hidden(X10,X9)|X9!=k3_tarski(X8))&(r2_hidden(esk3_3(X8,X9,X10),X8)|~r2_hidden(X10,X9)|X9!=k3_tarski(X8)))&(~r2_hidden(X12,X13)|~r2_hidden(X13,X8)|r2_hidden(X12,X9)|X9!=k3_tarski(X8)))&((~r2_hidden(esk4_2(X14,X15),X15)|(~r2_hidden(esk4_2(X14,X15),X17)|~r2_hidden(X17,X14))|X15=k3_tarski(X14))&((r2_hidden(esk4_2(X14,X15),esk5_2(X14,X15))|r2_hidden(esk4_2(X14,X15),X15)|X15=k3_tarski(X14))&(r2_hidden(esk5_2(X14,X15),X14)|r2_hidden(esk4_2(X14,X15),X15)|X15=k3_tarski(X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_tarski])])])])])])).
fof(c_0_32, plain, ![X23, X24]:((~r1_ordinal1(X23,X24)|r1_tarski(X23,X24)|(~v3_ordinal1(X23)|~v3_ordinal1(X24)))&(~r1_tarski(X23,X24)|r1_ordinal1(X23,X24)|(~v3_ordinal1(X23)|~v3_ordinal1(X24)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r1_ordinal1])])])).
cnf(c_0_33, plain, (r1_ordinal1(esk6_1(X1),X1)|v1_ordinal1(X1)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])).
cnf(c_0_34, negated_conjecture, (v3_ordinal1(X1)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_25, c_0_30])).
cnf(c_0_35, plain, (r1_tarski(X2,X1)|~v1_ordinal1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_36, plain, (r2_hidden(esk5_2(X1,X2),X1)|r2_hidden(esk4_2(X1,X2),X2)|X2=k3_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_37, plain, (r1_tarski(X1,X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, negated_conjecture, (r1_ordinal1(esk6_1(esk1_0),esk1_0)|v1_ordinal1(esk1_0)), inference(spm,[status(thm)],[c_0_33, c_0_30])).
cnf(c_0_39, negated_conjecture, (v1_ordinal1(esk1_0)|v3_ordinal1(esk6_1(esk1_0))), inference(spm,[status(thm)],[c_0_34, c_0_24])).
cnf(c_0_40, plain, (v1_ordinal1(X1)|~r1_tarski(esk6_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_41, plain, (X1=k3_tarski(X2)|r1_tarski(esk4_2(X2,X1),X1)|r2_hidden(esk5_2(X2,X1),X2)|~v1_ordinal1(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_42, negated_conjecture, (v1_ordinal1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_30])]), c_0_39]), c_0_40])).
cnf(c_0_43, plain, (X1=k3_tarski(X2)|r2_hidden(esk5_2(X2,X1),X2)|v3_ordinal1(esk4_2(X2,X1))|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_25, c_0_36])).
fof(c_0_44, plain, ![X33, X34, X35, X36, X37]:((~r1_tarski(X33,X34)|(~r2_hidden(X35,X33)|r2_hidden(X35,X34)))&((r2_hidden(esk7_2(X36,X37),X36)|r1_tarski(X36,X37))&(~r2_hidden(esk7_2(X36,X37),X37)|r1_tarski(X36,X37)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_45, negated_conjecture, (k3_tarski(X1)=esk1_0|r1_tarski(esk4_2(X1,esk1_0),esk1_0)|r2_hidden(esk5_2(X1,esk1_0),X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_46, negated_conjecture, (k3_tarski(X1)=esk1_0|r2_hidden(esk5_2(X1,esk1_0),X1)|v3_ordinal1(esk4_2(X1,esk1_0))), inference(spm,[status(thm)],[c_0_43, c_0_30])).
cnf(c_0_47, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_48, plain, (r2_hidden(esk4_2(X1,X2),esk5_2(X1,X2))|r2_hidden(esk4_2(X1,X2),X2)|X2=k3_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_49, negated_conjecture, (k3_tarski(X1)=esk1_0|r1_tarski(esk4_2(X1,esk1_0),esk1_0)|r1_tarski(esk5_2(X1,esk1_0),X1)|~v1_ordinal1(X1)), inference(spm,[status(thm)],[c_0_35, c_0_45])).
cnf(c_0_50, negated_conjecture, (k3_tarski(esk1_0)=esk1_0|v3_ordinal1(esk4_2(esk1_0,esk1_0))|v3_ordinal1(esk5_2(esk1_0,esk1_0))), inference(spm,[status(thm)],[c_0_34, c_0_46])).
cnf(c_0_51, plain, (X1=k3_tarski(X2)|r2_hidden(esk4_2(X2,X1),X1)|r2_hidden(esk4_2(X2,X1),X3)|~r1_tarski(esk5_2(X2,X1),X3)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_52, negated_conjecture, (k3_tarski(esk1_0)=esk1_0|r1_tarski(esk5_2(esk1_0,esk1_0),esk1_0)|r1_tarski(esk4_2(esk1_0,esk1_0),esk1_0)), inference(spm,[status(thm)],[c_0_49, c_0_42])).
cnf(c_0_53, negated_conjecture, (r1_tarski(X1,esk1_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_35, c_0_42])).
cnf(c_0_54, negated_conjecture, (k3_tarski(esk1_0)=esk1_0|v3_ordinal1(esk4_2(esk1_0,esk1_0))|v3_ordinal1(X1)|~r2_hidden(X1,esk5_2(esk1_0,esk1_0))), inference(spm,[status(thm)],[c_0_25, c_0_50])).
cnf(c_0_55, plain, (r2_hidden(X1,k1_ordinal1(X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_56, plain, (r1_ordinal1(X1,X2)|~r1_tarski(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_57, negated_conjecture, (k3_tarski(esk1_0)=esk1_0|r1_tarski(esk4_2(esk1_0,esk1_0),esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53])).
cnf(c_0_58, negated_conjecture, (k3_tarski(esk1_0)=esk1_0|v3_ordinal1(esk4_2(esk1_0,esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_48]), c_0_34])).
cnf(c_0_59, plain, (r2_hidden(X1,X2)|X1=X2|~r2_hidden(X1,k1_ordinal1(X2))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_60, plain, (r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2)), inference(rw,[status(thm)],[c_0_55, c_0_18])).
cnf(c_0_61, negated_conjecture, (k3_tarski(esk1_0)=esk1_0|r1_ordinal1(esk4_2(esk1_0,esk1_0),esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_30])]), c_0_58])).
cnf(c_0_62, plain, (X1=X2|r2_hidden(X1,X2)|~r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))), inference(rw,[status(thm)],[c_0_59, c_0_18])).
cnf(c_0_63, negated_conjecture, (k3_tarski(esk1_0)=esk1_0|r2_hidden(esk4_2(esk1_0,esk1_0),k2_xboole_0(esk1_0,k1_tarski(esk1_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_30])]), c_0_58])).
fof(c_0_64, plain, ![X44]:r2_hidden(X44,k1_ordinal1(X44)), inference(variable_rename,[status(thm)],[t10_ordinal1])).
cnf(c_0_65, negated_conjecture, (v4_ordinal1(esk1_0)|r2_hidden(k1_ordinal1(X1),esk1_0)|~v3_ordinal1(X1)|~r2_hidden(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_66, plain, ![X39]:((~v4_ordinal1(X39)|X39=k3_tarski(X39))&(X39!=k3_tarski(X39)|v4_ordinal1(X39))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_ordinal1])])).
cnf(c_0_67, plain, (X2=k3_tarski(X1)|~r2_hidden(esk4_2(X1,X2),X2)|~r2_hidden(esk4_2(X1,X2),X3)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_68, negated_conjecture, (esk4_2(esk1_0,esk1_0)=esk1_0|k3_tarski(esk1_0)=esk1_0|r2_hidden(esk4_2(esk1_0,esk1_0),esk1_0)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_69, plain, (r2_hidden(X1,k1_ordinal1(X1))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_70, negated_conjecture, (v4_ordinal1(esk1_0)|r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),esk1_0)|~v3_ordinal1(X1)|~r2_hidden(X1,esk1_0)), inference(rw,[status(thm)],[c_0_65, c_0_18])).
cnf(c_0_71, plain, (X1=k3_tarski(X1)|~v4_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_72, negated_conjecture, (esk4_2(esk1_0,esk1_0)=esk1_0|k3_tarski(esk1_0)=esk1_0|~r2_hidden(esk4_2(esk1_0,esk1_0),X1)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_73, plain, (r2_hidden(X1,k2_xboole_0(X1,k1_tarski(X1)))), inference(rw,[status(thm)],[c_0_69, c_0_18])).
fof(c_0_74, plain, ![X25, X26]:(~r2_hidden(X25,X26)|~r1_tarski(X26,X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_75, plain, (r1_tarski(X1,X2)|~r2_hidden(esk7_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_76, plain, (r2_hidden(esk7_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_77, negated_conjecture, (k3_tarski(esk1_0)=esk1_0|r2_hidden(k2_xboole_0(esk4_2(esk1_0,esk1_0),k1_tarski(esk4_2(esk1_0,esk1_0))),esk1_0)|~r2_hidden(esk4_2(esk1_0,esk1_0),esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_58]), c_0_71])).
cnf(c_0_78, negated_conjecture, (esk4_2(esk1_0,esk1_0)=esk1_0|k3_tarski(esk1_0)=esk1_0|~r2_hidden(k2_xboole_0(esk4_2(esk1_0,esk1_0),k1_tarski(esk4_2(esk1_0,esk1_0))),esk1_0)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_79, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_80, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_81, negated_conjecture, (X1=k3_tarski(esk1_0)|r1_tarski(esk5_2(esk1_0,X1),esk1_0)|r2_hidden(esk4_2(esk1_0,X1),X1)), inference(spm,[status(thm)],[c_0_53, c_0_36])).
cnf(c_0_82, negated_conjecture, (esk4_2(esk1_0,esk1_0)=esk1_0|k3_tarski(esk1_0)=esk1_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_68]), c_0_78])).
cnf(c_0_83, plain, (~r2_hidden(X1,X1)), inference(spm,[status(thm)],[c_0_79, c_0_80])).
fof(c_0_84, plain, ![X27, X28]:((~r2_hidden(X27,X28)|r1_ordinal1(k1_ordinal1(X27),X28)|~v3_ordinal1(X28)|~v3_ordinal1(X27))&(~r1_ordinal1(k1_ordinal1(X27),X28)|r2_hidden(X27,X28)|~v3_ordinal1(X28)|~v3_ordinal1(X27))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t33_ordinal1])])])])).
cnf(c_0_85, negated_conjecture, (X1=k3_tarski(esk1_0)|r2_hidden(esk4_2(esk1_0,X1),X1)|r2_hidden(X2,esk1_0)|~r2_hidden(X2,esk5_2(esk1_0,X1))), inference(spm,[status(thm)],[c_0_47, c_0_81])).
cnf(c_0_86, negated_conjecture, (k3_tarski(esk1_0)=esk1_0|r2_hidden(esk1_0,esk5_2(esk1_0,esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_82]), c_0_83])).
cnf(c_0_87, plain, (r1_ordinal1(k1_ordinal1(X1),X2)|~r2_hidden(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_88, negated_conjecture, (k3_tarski(esk1_0)=esk1_0|r2_hidden(esk4_2(esk1_0,esk1_0),esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_83])).
cnf(c_0_89, plain, (r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_87, c_0_18])).
cnf(c_0_90, plain, (v4_ordinal1(X1)|X1!=k3_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_91, negated_conjecture, (k3_tarski(esk1_0)=esk1_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_82]), c_0_83])).
fof(c_0_92, plain, ![X32]:(~v3_ordinal1(X32)|v3_ordinal1(k1_ordinal1(X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t29_ordinal1])])).
cnf(c_0_93, plain, (r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X2)|~r2_hidden(X1,X2)|~v3_ordinal1(X2)), inference(csr,[status(thm)],[c_0_89, c_0_25])).
cnf(c_0_94, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v4_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_95, negated_conjecture, (v4_ordinal1(esk1_0)), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_96, plain, (v3_ordinal1(k1_ordinal1(X1))|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_97, negated_conjecture, (v3_ordinal1(esk2_0)|~v4_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_98, negated_conjecture, (r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),esk1_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_93, c_0_30])).
cnf(c_0_99, negated_conjecture, (r2_hidden(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_94, c_0_95])])).
cnf(c_0_100, plain, (v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)), inference(rw,[status(thm)],[c_0_96, c_0_18])).
cnf(c_0_101, negated_conjecture, (v3_ordinal1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_95])])).
cnf(c_0_102, negated_conjecture, (~r2_hidden(k1_ordinal1(esk2_0),esk1_0)|~v4_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_103, negated_conjecture, (r1_ordinal1(k2_xboole_0(esk2_0,k1_tarski(esk2_0)),esk1_0)), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_104, negated_conjecture, (v3_ordinal1(k2_xboole_0(esk2_0,k1_tarski(esk2_0)))), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_105, negated_conjecture, (~v4_ordinal1(esk1_0)|~r2_hidden(k2_xboole_0(esk2_0,k1_tarski(esk2_0)),esk1_0)), inference(rw,[status(thm)],[c_0_102, c_0_18])).
cnf(c_0_106, plain, (r2_hidden(esk3_3(X1,X2,X3),X1)|~r2_hidden(X3,X2)|X2!=k3_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_107, negated_conjecture, (r2_hidden(k2_xboole_0(esk2_0,k1_tarski(esk2_0)),k2_xboole_0(esk1_0,k1_tarski(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_103]), c_0_30]), c_0_104])])).
cnf(c_0_108, negated_conjecture, (~r2_hidden(k2_xboole_0(esk2_0,k1_tarski(esk2_0)),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_105, c_0_95])])).
cnf(c_0_109, plain, (r2_hidden(esk3_3(X1,k2_xboole_0(X2,k1_tarski(X2)),X2),X1)|k2_xboole_0(X2,k1_tarski(X2))!=k3_tarski(X1)), inference(spm,[status(thm)],[c_0_106, c_0_73])).
cnf(c_0_110, negated_conjecture, (k2_xboole_0(esk2_0,k1_tarski(esk2_0))=esk1_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_107]), c_0_108])).
cnf(c_0_111, negated_conjecture, (r2_hidden(esk3_3(X1,esk1_0,esk2_0),X1)|k3_tarski(X1)!=esk1_0), inference(spm,[status(thm)],[c_0_109, c_0_110])).
cnf(c_0_112, plain, (r2_hidden(X1,esk3_3(X2,X3,X1))|~r2_hidden(X1,X3)|X3!=k3_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_113, negated_conjecture, (r1_ordinal1(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_110]), c_0_101])]), c_0_34])).
cnf(c_0_114, negated_conjecture, (r2_hidden(esk3_3(esk1_0,esk1_0,esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_111, c_0_91])).
cnf(c_0_115, plain, (r2_hidden(X1,esk3_3(X2,k3_tarski(X2),X1))|~r2_hidden(X1,k3_tarski(X2))), inference(er,[status(thm)],[c_0_112])).
cnf(c_0_116, negated_conjecture, (r1_ordinal1(esk3_3(esk1_0,esk1_0,esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_113, c_0_114])).
cnf(c_0_117, negated_conjecture, (v3_ordinal1(esk3_3(esk1_0,esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_114]), c_0_30])])).
cnf(c_0_118, negated_conjecture, (r2_hidden(X1,esk3_3(esk1_0,esk1_0,X1))|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_115, c_0_91])).
cnf(c_0_119, negated_conjecture, (r1_tarski(esk3_3(esk1_0,esk1_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_116]), c_0_101]), c_0_117])])).
cnf(c_0_120, negated_conjecture, (r2_hidden(esk2_0,esk3_3(esk1_0,esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_118, c_0_99])).
cnf(c_0_121, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_119]), c_0_120])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 122
# Proof object clause steps            : 93
# Proof object formula steps           : 29
# Proof object conjectures             : 53
# Proof object clause conjectures      : 50
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 30
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 50
# Proof object simplifying inferences  : 48
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 32
# Processed clauses                    : 1426
# ...of these trivial                  : 30
# ...subsumed                          : 586
# ...remaining for further processing  : 810
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 51
# Backward-rewritten                   : 231
# Generated clauses                    : 14973
# ...of the previous two non-trivial   : 14702
# Contextual simplify-reflections      : 52
# Paramodulations                      : 14888
# Factorizations                       : 0
# Equation resolutions                 : 85
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 527
#    Positive orientable unit clauses  : 53
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 11
#    Non-unit-clauses                  : 463
# Current number of unprocessed clauses: 13198
# ...number of literals in the above   : 56020
# Current number of archived formulas  : 0
# Current number of archived clauses   : 283
# Clause-clause subsumption calls (NU) : 46742
# Rec. Clause-clause subsumption calls : 21795
# Non-unit clause-clause subsumptions  : 609
# Unit Clause-clause subsumption calls : 2511
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 71
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 416622

# -------------------------------------------------
# User time                : 0.196 s
# System time              : 0.010 s
# Total time               : 0.206 s
# Maximum resident set size: 28756 pages
