# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t37_wellord1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>((v2_wellord1(X3)&(r2_hidden(X1,k3_relat_1(X3))&r2_hidden(X2,k3_relat_1(X3))))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>r1_tarski(k1_wellord1(X3,X1),k1_wellord1(X3,X2))))), file('data/MPTP/proofs/2db0ba608b834704b1c8bb0f71faf936.E_input', t37_wellord1)).
fof(l4_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v6_relat_2(X1)<=>![X2, X3]:~((r2_hidden(X2,k3_relat_1(X1))&(r2_hidden(X3,k3_relat_1(X1))&(~(X2=X3)&(~(r2_hidden(k4_tarski(X2,X3),X1))&~(r2_hidden(k4_tarski(X3,X2),X1))))))))), file('data/MPTP/proofs/2db0ba608b834704b1c8bb0f71faf936.E_input', l4_wellord1)).
fof(d4_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v2_wellord1(X1)<=>(v1_relat_2(X1)&(v8_relat_2(X1)&(v4_relat_2(X1)&(v6_relat_2(X1)&v1_wellord1(X1))))))), file('data/MPTP/proofs/2db0ba608b834704b1c8bb0f71faf936.E_input', d4_wellord1)).
fof(d1_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(X3=k1_wellord1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(~(X4=X2)&r2_hidden(k4_tarski(X4,X2),X1))))), file('data/MPTP/proofs/2db0ba608b834704b1c8bb0f71faf936.E_input', d1_wellord1)).
fof(l2_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v8_relat_2(X1)<=>![X2, X3, X4]:((r2_hidden(k4_tarski(X2,X3),X1)&r2_hidden(k4_tarski(X3,X4),X1))=>r2_hidden(k4_tarski(X2,X4),X1)))), file('data/MPTP/proofs/2db0ba608b834704b1c8bb0f71faf936.E_input', l2_wellord1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/2db0ba608b834704b1c8bb0f71faf936.E_input', d3_tarski)).
fof(s1_xboole_0__e7_18__wellord1, axiom, ![X1, X2]:(v1_relat_1(X1)=>?[X3]:![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,k3_relat_1(X1))&~(r2_hidden(k4_tarski(X4,X2),X1))))), file('data/MPTP/proofs/2db0ba608b834704b1c8bb0f71faf936.E_input', s1_xboole_0__e7_18__wellord1)).
fof(l3_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v4_relat_2(X1)<=>![X2, X3]:((r2_hidden(k4_tarski(X2,X3),X1)&r2_hidden(k4_tarski(X3,X2),X1))=>X2=X3))), file('data/MPTP/proofs/2db0ba608b834704b1c8bb0f71faf936.E_input', l3_wellord1)).
fof(t10_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v2_wellord1(X1)=>![X2]:~((r1_tarski(X2,k3_relat_1(X1))&(~(X2=k1_xboole_0)&![X3]:~((r2_hidden(X3,X2)&![X4]:(r2_hidden(X4,X2)=>r2_hidden(k4_tarski(X3,X4),X1))))))))), file('data/MPTP/proofs/2db0ba608b834704b1c8bb0f71faf936.E_input', t10_wellord1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/2db0ba608b834704b1c8bb0f71faf936.E_input', d10_xboole_0)).
fof(t20_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)=>(r2_hidden(X1,k1_relat_1(X3))&r2_hidden(X2,k2_relat_1(X3))))), file('data/MPTP/proofs/2db0ba608b834704b1c8bb0f71faf936.E_input', t20_relat_1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/2db0ba608b834704b1c8bb0f71faf936.E_input', t7_ordinal1)).
fof(c_0_12, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>((v2_wellord1(X3)&(r2_hidden(X1,k3_relat_1(X3))&r2_hidden(X2,k3_relat_1(X3))))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>r1_tarski(k1_wellord1(X3,X1),k1_wellord1(X3,X2)))))), inference(assume_negation,[status(cth)],[t37_wellord1])).
fof(c_0_13, plain, ![X8, X9, X10]:((~v6_relat_2(X8)|(~r2_hidden(X9,k3_relat_1(X8))|(~r2_hidden(X10,k3_relat_1(X8))|(X9=X10|(r2_hidden(k4_tarski(X9,X10),X8)|r2_hidden(k4_tarski(X10,X9),X8)))))|~v1_relat_1(X8))&((r2_hidden(esk4_1(X8),k3_relat_1(X8))|v6_relat_2(X8)|~v1_relat_1(X8))&((r2_hidden(esk5_1(X8),k3_relat_1(X8))|v6_relat_2(X8)|~v1_relat_1(X8))&((esk4_1(X8)!=esk5_1(X8)|v6_relat_2(X8)|~v1_relat_1(X8))&((~r2_hidden(k4_tarski(esk4_1(X8),esk5_1(X8)),X8)|v6_relat_2(X8)|~v1_relat_1(X8))&(~r2_hidden(k4_tarski(esk5_1(X8),esk4_1(X8)),X8)|v6_relat_2(X8)|~v1_relat_1(X8))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l4_wellord1])])])])])])])).
fof(c_0_14, negated_conjecture, (v1_relat_1(esk3_0)&((v2_wellord1(esk3_0)&(r2_hidden(esk1_0,k3_relat_1(esk3_0))&r2_hidden(esk2_0,k3_relat_1(esk3_0))))&((~r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|~r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0)))&(r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
cnf(c_0_15, plain, (X2=X3|r2_hidden(k4_tarski(X2,X3),X1)|r2_hidden(k4_tarski(X3,X2),X1)|~v6_relat_2(X1)|~r2_hidden(X2,k3_relat_1(X1))|~r2_hidden(X3,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_16, negated_conjecture, (r2_hidden(esk2_0,k3_relat_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_17, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, plain, ![X13]:(((v1_relat_2(X13)|~v2_wellord1(X13)|~v1_relat_1(X13))&((v8_relat_2(X13)|~v2_wellord1(X13)|~v1_relat_1(X13))&((v4_relat_2(X13)|~v2_wellord1(X13)|~v1_relat_1(X13))&((v6_relat_2(X13)|~v2_wellord1(X13)|~v1_relat_1(X13))&(v1_wellord1(X13)|~v2_wellord1(X13)|~v1_relat_1(X13))))))&(~v1_relat_2(X13)|(~v8_relat_2(X13)|(~v4_relat_2(X13)|(~v6_relat_2(X13)|~v1_wellord1(X13))))|v2_wellord1(X13)|~v1_relat_1(X13))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_wellord1])])])).
fof(c_0_19, plain, ![X14, X15, X16, X17, X18, X19, X20]:((((X17!=X15|~r2_hidden(X17,X16)|X16!=k1_wellord1(X14,X15)|~v1_relat_1(X14))&(r2_hidden(k4_tarski(X17,X15),X14)|~r2_hidden(X17,X16)|X16!=k1_wellord1(X14,X15)|~v1_relat_1(X14)))&(X18=X15|~r2_hidden(k4_tarski(X18,X15),X14)|r2_hidden(X18,X16)|X16!=k1_wellord1(X14,X15)|~v1_relat_1(X14)))&((~r2_hidden(esk6_3(X14,X19,X20),X20)|(esk6_3(X14,X19,X20)=X19|~r2_hidden(k4_tarski(esk6_3(X14,X19,X20),X19),X14))|X20=k1_wellord1(X14,X19)|~v1_relat_1(X14))&((esk6_3(X14,X19,X20)!=X19|r2_hidden(esk6_3(X14,X19,X20),X20)|X20=k1_wellord1(X14,X19)|~v1_relat_1(X14))&(r2_hidden(k4_tarski(esk6_3(X14,X19,X20),X19),X14)|r2_hidden(esk6_3(X14,X19,X20),X20)|X20=k1_wellord1(X14,X19)|~v1_relat_1(X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_wellord1])])])])])])])).
cnf(c_0_20, negated_conjecture, (X1=esk2_0|r2_hidden(k4_tarski(esk2_0,X1),esk3_0)|r2_hidden(k4_tarski(X1,esk2_0),esk3_0)|~v6_relat_2(esk3_0)|~r2_hidden(X1,k3_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])])).
cnf(c_0_21, plain, (v6_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, negated_conjecture, (v2_wellord1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23, plain, (X1=X2|r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_wellord1(X3,X2)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, negated_conjecture, (X1=esk2_0|r2_hidden(k4_tarski(X1,esk2_0),esk3_0)|r2_hidden(k4_tarski(esk2_0,X1),esk3_0)|~r2_hidden(X1,k3_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_17])])).
cnf(c_0_25, negated_conjecture, (r2_hidden(esk1_0,k3_relat_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_26, plain, ![X22, X23, X24, X25]:((~v8_relat_2(X22)|(~r2_hidden(k4_tarski(X23,X24),X22)|~r2_hidden(k4_tarski(X24,X25),X22)|r2_hidden(k4_tarski(X23,X25),X22))|~v1_relat_1(X22))&(((r2_hidden(k4_tarski(esk7_1(X22),esk8_1(X22)),X22)|v8_relat_2(X22)|~v1_relat_1(X22))&(r2_hidden(k4_tarski(esk8_1(X22),esk9_1(X22)),X22)|v8_relat_2(X22)|~v1_relat_1(X22)))&(~r2_hidden(k4_tarski(esk7_1(X22),esk9_1(X22)),X22)|v8_relat_2(X22)|~v1_relat_1(X22)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l2_wellord1])])])])])).
cnf(c_0_27, plain, (X1=X2|r2_hidden(X1,k1_wellord1(X3,X2))|~r2_hidden(k4_tarski(X1,X2),X3)|~v1_relat_1(X3)), inference(er,[status(thm)],[c_0_23])).
cnf(c_0_28, negated_conjecture, (esk2_0=esk1_0|r2_hidden(k4_tarski(esk2_0,esk1_0),esk3_0)|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,X4)|X4!=k1_wellord1(X3,X2)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_30, plain, ![X29, X30, X31, X32, X33]:((~r1_tarski(X29,X30)|(~r2_hidden(X31,X29)|r2_hidden(X31,X30)))&((r2_hidden(esk10_2(X32,X33),X32)|r1_tarski(X32,X33))&(~r2_hidden(esk10_2(X32,X33),X33)|r1_tarski(X32,X33)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_31, plain, ![X35, X36, X38, X39]:(((r2_hidden(X38,k3_relat_1(X35))|~r2_hidden(X38,esk11_2(X35,X36))|~v1_relat_1(X35))&(~r2_hidden(k4_tarski(X38,X36),X35)|~r2_hidden(X38,esk11_2(X35,X36))|~v1_relat_1(X35)))&(~r2_hidden(X39,k3_relat_1(X35))|r2_hidden(k4_tarski(X39,X36),X35)|r2_hidden(X39,esk11_2(X35,X36))|~v1_relat_1(X35))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[s1_xboole_0__e7_18__wellord1])])])])])])])).
cnf(c_0_32, plain, (r2_hidden(k4_tarski(X2,X4),X1)|~v8_relat_2(X1)|~r2_hidden(k4_tarski(X2,X3),X1)|~r2_hidden(k4_tarski(X3,X4),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, negated_conjecture, (esk2_0=esk1_0|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_17])])).
cnf(c_0_34, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,k1_wellord1(X3,X2))|~v1_relat_1(X3)), inference(er,[status(thm)],[c_0_29])).
cnf(c_0_35, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_36, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_37, plain, (~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,esk11_2(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))|r2_hidden(k4_tarski(X1,esk2_0),esk3_0)|~v8_relat_2(esk3_0)|~r2_hidden(k4_tarski(X1,esk1_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_17])])).
cnf(c_0_39, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~v8_relat_2(X3)|~r2_hidden(k4_tarski(X1,X4),X3)|~r2_hidden(X4,k1_wellord1(X3,X2))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_32, c_0_34])).
cnf(c_0_40, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r2_hidden(X1,k1_wellord1(esk3_0,esk2_0))|~r2_hidden(X1,k1_wellord1(esk3_0,esk1_0))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_41, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))|~r2_hidden(esk1_0,esk11_2(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_33]), c_0_17])])).
cnf(c_0_42, negated_conjecture, (esk2_0=esk1_0|X1=esk2_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))|r2_hidden(X1,k1_wellord1(esk3_0,esk2_0))|~v8_relat_2(esk3_0)|~r2_hidden(k4_tarski(X1,esk1_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_38]), c_0_17])])).
cnf(c_0_43, negated_conjecture, (esk2_0=esk1_0|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r2_hidden(k4_tarski(esk2_0,X1),esk3_0)|~v8_relat_2(esk3_0)|~r2_hidden(esk1_0,k1_wellord1(esk3_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_28]), c_0_17])])).
cnf(c_0_44, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))|r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_33]), c_0_17])])).
fof(c_0_45, plain, ![X40, X41, X42]:((~v4_relat_2(X40)|(~r2_hidden(k4_tarski(X41,X42),X40)|~r2_hidden(k4_tarski(X42,X41),X40)|X41=X42)|~v1_relat_1(X40))&(((r2_hidden(k4_tarski(esk12_1(X40),esk13_1(X40)),X40)|v4_relat_2(X40)|~v1_relat_1(X40))&(r2_hidden(k4_tarski(esk13_1(X40),esk12_1(X40)),X40)|v4_relat_2(X40)|~v1_relat_1(X40)))&(esk12_1(X40)!=esk13_1(X40)|v4_relat_2(X40)|~v1_relat_1(X40)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_wellord1])])])])])).
cnf(c_0_46, plain, (X1!=X2|~r2_hidden(X1,X3)|X3!=k1_wellord1(X4,X2)|~v1_relat_1(X4)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_47, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk2_0))|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|~r2_hidden(esk1_0,esk11_2(esk3_0,esk2_0))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_48, negated_conjecture, (esk2_0=esk1_0|X1=esk2_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))|r2_hidden(X1,k1_wellord1(esk3_0,esk2_0))|~v8_relat_2(esk3_0)|~r2_hidden(X1,k1_wellord1(esk3_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_34]), c_0_17])])).
cnf(c_0_49, plain, (v8_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_50, negated_conjecture, (esk2_0=esk1_0|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r2_hidden(k4_tarski(X1,X2),esk3_0)|~v8_relat_2(esk3_0)|~r2_hidden(k4_tarski(X1,esk2_0),esk3_0)|~r2_hidden(esk1_0,k1_wellord1(esk3_0,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_43]), c_0_17])])).
cnf(c_0_51, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk2_0))|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_40, c_0_44])).
cnf(c_0_52, plain, (X2=X3|~v4_relat_2(X1)|~r2_hidden(k4_tarski(X2,X3),X1)|~r2_hidden(k4_tarski(X3,X2),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_53, plain, (r2_hidden(k4_tarski(X1,X3),X2)|r2_hidden(X1,esk11_2(X2,X3))|~r2_hidden(X1,k3_relat_1(X2))|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_54, plain, (~r2_hidden(X1,k1_wellord1(X2,X1))|~v1_relat_1(X2)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_46])])).
cnf(c_0_55, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk2_0))|~r2_hidden(esk1_0,esk11_2(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_47]), c_0_17])])).
cnf(c_0_56, negated_conjecture, (esk2_0=esk1_0|X1=esk2_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))|r2_hidden(X1,k1_wellord1(esk3_0,esk2_0))|~r2_hidden(X1,k1_wellord1(esk3_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_22]), c_0_17])])).
cnf(c_0_57, plain, (r2_hidden(esk10_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_58, negated_conjecture, (esk2_0=esk1_0|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r2_hidden(k4_tarski(X1,X2),esk3_0)|~v8_relat_2(esk3_0)|~r2_hidden(esk1_0,k1_wellord1(esk3_0,X2))|~r2_hidden(X1,k1_wellord1(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_34]), c_0_17])])).
cnf(c_0_59, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk2_0))|r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_51]), c_0_17])])).
cnf(c_0_60, plain, (X1=X2|r2_hidden(X2,esk11_2(X3,X1))|~v4_relat_2(X3)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X2,k3_relat_1(X3))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_61, negated_conjecture, (esk2_0=esk1_0|~r2_hidden(esk1_0,esk11_2(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_17])])).
cnf(c_0_62, plain, (r1_tarski(X1,X2)|~r2_hidden(esk10_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_63, negated_conjecture, (esk10_2(k1_wellord1(esk3_0,esk1_0),X1)=esk2_0|esk2_0=esk1_0|r1_tarski(k1_wellord1(esk3_0,esk1_0),X1)|r2_hidden(esk10_2(k1_wellord1(esk3_0,esk1_0),X1),k1_wellord1(esk3_0,esk2_0))|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_64, negated_conjecture, (esk2_0=esk1_0|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r2_hidden(k4_tarski(X1,X2),esk3_0)|~r2_hidden(esk1_0,k1_wellord1(esk3_0,X2))|~r2_hidden(X1,k1_wellord1(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_49]), c_0_22]), c_0_17])])).
cnf(c_0_65, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_59]), c_0_17])])).
cnf(c_0_66, negated_conjecture, (esk2_0=esk1_0|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|~v4_relat_2(esk3_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_28]), c_0_25]), c_0_17])]), c_0_61])).
cnf(c_0_67, negated_conjecture, (esk10_2(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))=esk2_0|esk2_0=esk1_0|r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_68, negated_conjecture, (esk2_0=esk1_0|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r2_hidden(k4_tarski(X1,esk2_0),esk3_0)|~r2_hidden(X1,k1_wellord1(esk3_0,esk2_0))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_69, negated_conjecture, (esk2_0=esk1_0|~v4_relat_2(esk3_0)|~r2_hidden(k4_tarski(esk2_0,esk1_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_66]), c_0_17])])).
cnf(c_0_70, negated_conjecture, (~r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|~r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_71, negated_conjecture, (esk2_0=esk1_0|r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))), inference(spm,[status(thm)],[c_0_57, c_0_67])).
cnf(c_0_72, negated_conjecture, (esk2_0=esk1_0|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_68, c_0_65])).
fof(c_0_73, plain, ![X48, X49, X51]:((r2_hidden(esk14_2(X48,X49),X49)|X49=k1_xboole_0|~r1_tarski(X49,k3_relat_1(X48))|~v2_wellord1(X48)|~v1_relat_1(X48))&(~r2_hidden(X51,X49)|r2_hidden(k4_tarski(esk14_2(X48,X49),X51),X48)|X49=k1_xboole_0|~r1_tarski(X49,k3_relat_1(X48))|~v2_wellord1(X48)|~v1_relat_1(X48))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t10_wellord1])])])])])])).
cnf(c_0_74, negated_conjecture, (esk2_0=esk1_0|~v4_relat_2(esk3_0)|~r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_34]), c_0_17])])).
cnf(c_0_75, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_72])).
fof(c_0_76, plain, ![X52, X53]:(((r1_tarski(X52,X53)|X52!=X53)&(r1_tarski(X53,X52)|X52!=X53))&(~r1_tarski(X52,X53)|~r1_tarski(X53,X52)|X52=X53)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_77, plain, (r2_hidden(k4_tarski(esk14_2(X3,X2),X1),X3)|X2=k1_xboole_0|~r2_hidden(X1,X2)|~r1_tarski(X2,k3_relat_1(X3))|~v2_wellord1(X3)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_78, plain, (r2_hidden(X1,k3_relat_1(X2))|~r2_hidden(X1,esk11_2(X2,X3))|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_79, negated_conjecture, (esk2_0=esk1_0|~v4_relat_2(esk3_0)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_80, plain, (v4_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_81, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_82, plain, (X1=k1_xboole_0|~r1_tarski(X1,k3_relat_1(X2))|~r2_hidden(esk14_2(X2,X1),esk11_2(X2,X3))|~r2_hidden(X3,X1)|~v2_wellord1(X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_37, c_0_77])).
cnf(c_0_83, plain, (r2_hidden(esk14_2(X1,X2),X2)|X2=k1_xboole_0|~r1_tarski(X2,k3_relat_1(X1))|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_84, plain, (r1_tarski(esk11_2(X1,X2),X3)|r2_hidden(esk10_2(esk11_2(X1,X2),X3),k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_78, c_0_57])).
cnf(c_0_85, negated_conjecture, (esk2_0=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_22]), c_0_17])])).
cnf(c_0_86, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_81])).
fof(c_0_87, plain, ![X45, X46, X47]:((r2_hidden(X45,k1_relat_1(X47))|~r2_hidden(k4_tarski(X45,X46),X47)|~v1_relat_1(X47))&(r2_hidden(X46,k2_relat_1(X47))|~r2_hidden(k4_tarski(X45,X46),X47)|~v1_relat_1(X47))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t20_relat_1])])])).
cnf(c_0_88, plain, (esk11_2(X1,X2)=k1_xboole_0|~r1_tarski(esk11_2(X1,X2),k3_relat_1(X1))|~r2_hidden(X2,esk11_2(X1,X2))|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_89, plain, (r1_tarski(esk11_2(X1,X2),k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_62, c_0_84])).
cnf(c_0_90, negated_conjecture, (~r2_hidden(k4_tarski(esk1_0,esk1_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_85]), c_0_86]), c_0_85])])).
cnf(c_0_91, plain, (r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_92, plain, (esk11_2(X1,X2)=k1_xboole_0|~r2_hidden(X2,esk11_2(X1,X2))|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_93, negated_conjecture, (r2_hidden(esk1_0,esk11_2(esk3_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_53]), c_0_25]), c_0_17])])).
cnf(c_0_94, plain, (r2_hidden(X1,esk11_2(X2,X3))|r2_hidden(X3,k2_relat_1(X2))|~r2_hidden(X1,k3_relat_1(X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_91, c_0_53])).
cnf(c_0_95, negated_conjecture, (esk11_2(esk3_0,esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_22]), c_0_17])])).
cnf(c_0_96, plain, (r1_tarski(X1,esk11_2(X2,X3))|r2_hidden(X3,k2_relat_1(X2))|~r2_hidden(esk10_2(X1,esk11_2(X2,X3)),k3_relat_1(X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_62, c_0_94])).
cnf(c_0_97, negated_conjecture, (r1_tarski(k1_xboole_0,X1)|r2_hidden(esk10_2(k1_xboole_0,X1),k3_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_95]), c_0_17])])).
cnf(c_0_98, plain, (r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(X3,k1_wellord1(X2,X1))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_91, c_0_34])).
cnf(c_0_99, negated_conjecture, (r1_tarski(k1_xboole_0,esk11_2(esk3_0,X1))|r2_hidden(X1,k2_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_17])])).
cnf(c_0_100, plain, (r1_tarski(k1_wellord1(X1,X2),X3)|r2_hidden(X2,k2_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_98, c_0_57])).
cnf(c_0_101, plain, (~r2_hidden(X1,esk11_2(X2,X3))|~r2_hidden(X1,k1_wellord1(X2,X3))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_37, c_0_34])).
cnf(c_0_102, negated_conjecture, (r2_hidden(X1,esk11_2(esk3_0,X2))|r2_hidden(X2,k2_relat_1(esk3_0))|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_35, c_0_99])).
cnf(c_0_103, plain, (k1_wellord1(X1,X2)=k1_xboole_0|r2_hidden(X2,k2_relat_1(X1))|~v2_wellord1(X3)|~v1_relat_1(X1)|~v1_relat_1(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_83]), c_0_100])).
fof(c_0_104, plain, ![X54, X55]:(~r2_hidden(X54,X55)|~r1_tarski(X55,X54)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_105, negated_conjecture, (r2_hidden(X1,k2_relat_1(esk3_0))|~r2_hidden(X2,k1_wellord1(esk3_0,X1))|~r2_hidden(X2,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_17])])).
cnf(c_0_106, negated_conjecture, (k1_wellord1(X1,X2)=k1_xboole_0|r2_hidden(X2,k2_relat_1(X1))|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_22]), c_0_17])])).
cnf(c_0_107, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_108, negated_conjecture, (r2_hidden(X1,k2_relat_1(esk3_0))|~r2_hidden(X2,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_17])])).
cnf(c_0_109, negated_conjecture, (r2_hidden(esk1_0,k1_xboole_0)), inference(rw,[status(thm)],[c_0_93, c_0_95])).
cnf(c_0_110, plain, (~r2_hidden(X1,X1)), inference(spm,[status(thm)],[c_0_107, c_0_86])).
cnf(c_0_111, negated_conjecture, (r2_hidden(X1,k2_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_108, c_0_109])).
cnf(c_0_112, negated_conjecture, ($false), inference(spm,[status(thm)],[c_0_110, c_0_111]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 113
# Proof object clause steps            : 88
# Proof object formula steps           : 25
# Proof object conjectures             : 53
# Proof object clause conjectures      : 50
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 56
# Proof object simplifying inferences  : 76
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 46
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 46
# Processed clauses                    : 3859
# ...of these trivial                  : 9
# ...subsumed                          : 2308
# ...remaining for further processing  : 1542
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 443
# Backward-rewritten                   : 473
# Generated clauses                    : 18121
# ...of the previous two non-trivial   : 16729
# Contextual simplify-reflections      : 25
# Paramodulations                      : 18104
# Factorizations                       : 12
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 621
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 606
# Current number of unprocessed clauses: 11157
# ...number of literals in the above   : 80666
# Current number of archived formulas  : 0
# Current number of archived clauses   : 916
# Clause-clause subsumption calls (NU) : 234136
# Rec. Clause-clause subsumption calls : 19624
# Non-unit clause-clause subsumptions  : 2710
# Unit Clause-clause subsumption calls : 1146
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 24
# BW rewrite match successes           : 18
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 478456

# -------------------------------------------------
# User time                : 0.449 s
# System time              : 0.000 s
# Total time               : 0.449 s
# Maximum resident set size: 28744 pages
