# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/2ef03678974c4f6e93b3dbf6a3d47c54.E_input', d4_xboole_0)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/2ef03678974c4f6e93b3dbf6a3d47c54.E_input', t48_xboole_1)).
fof(t52_zfmisc_1, conjecture, ![X1, X2]:(r2_hidden(X1,X2)=>k3_xboole_0(X2,k1_tarski(X1))=k1_tarski(X1)), file('data/MPTP/proofs/2ef03678974c4f6e93b3dbf6a3d47c54.E_input', t52_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/2ef03678974c4f6e93b3dbf6a3d47c54.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/2ef03678974c4f6e93b3dbf6a3d47c54.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/2ef03678974c4f6e93b3dbf6a3d47c54.E_input', t71_enumset1)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/2ef03678974c4f6e93b3dbf6a3d47c54.E_input', d2_tarski)).
fof(c_0_7, plain, ![X24, X25, X26, X27, X28, X29, X30, X31]:((((r2_hidden(X27,X24)|~r2_hidden(X27,X26)|X26!=k3_xboole_0(X24,X25))&(r2_hidden(X27,X25)|~r2_hidden(X27,X26)|X26!=k3_xboole_0(X24,X25)))&(~r2_hidden(X28,X24)|~r2_hidden(X28,X25)|r2_hidden(X28,X26)|X26!=k3_xboole_0(X24,X25)))&((~r2_hidden(esk4_3(X29,X30,X31),X31)|(~r2_hidden(esk4_3(X29,X30,X31),X29)|~r2_hidden(esk4_3(X29,X30,X31),X30))|X31=k3_xboole_0(X29,X30))&((r2_hidden(esk4_3(X29,X30,X31),X29)|r2_hidden(esk4_3(X29,X30,X31),X31)|X31=k3_xboole_0(X29,X30))&(r2_hidden(esk4_3(X29,X30,X31),X30)|r2_hidden(esk4_3(X29,X30,X31),X31)|X31=k3_xboole_0(X29,X30))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_8, plain, ![X13, X14]:k4_xboole_0(X13,k4_xboole_0(X13,X14))=k3_xboole_0(X13,X14), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_9, negated_conjecture, ~(![X1, X2]:(r2_hidden(X1,X2)=>k3_xboole_0(X2,k1_tarski(X1))=k1_tarski(X1))), inference(assume_negation,[status(cth)],[t52_zfmisc_1])).
cnf(c_0_10, plain, (r2_hidden(esk4_3(X1,X2,X3),X1)|r2_hidden(esk4_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_12, negated_conjecture, (r2_hidden(esk1_0,esk2_0)&k3_xboole_0(esk2_0,k1_tarski(esk1_0))!=k1_tarski(esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
fof(c_0_13, plain, ![X7]:k2_tarski(X7,X7)=k1_tarski(X7), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_14, plain, ![X11, X12]:k1_enumset1(X11,X11,X12)=k2_tarski(X11,X12), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_15, plain, ![X8, X9, X10]:k2_enumset1(X8,X8,X9,X10)=k1_enumset1(X8,X9,X10), inference(variable_rename,[status(thm)],[t71_enumset1])).
cnf(c_0_16, plain, (X3=k3_xboole_0(X1,X2)|~r2_hidden(esk4_3(X1,X2,X3),X3)|~r2_hidden(esk4_3(X1,X2,X3),X1)|~r2_hidden(esk4_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_17, plain, (X3=k4_xboole_0(X1,k4_xboole_0(X1,X2))|r2_hidden(esk4_3(X1,X2,X3),X3)|r2_hidden(esk4_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_18, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X4!=k3_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_19, negated_conjecture, (k3_xboole_0(esk2_0,k1_tarski(esk1_0))!=k1_tarski(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (r2_hidden(esk4_3(X1,X2,X3),X2)|r2_hidden(esk4_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_24, plain, (X3=k4_xboole_0(X1,k4_xboole_0(X1,X2))|~r2_hidden(esk4_3(X1,X2,X3),X3)|~r2_hidden(esk4_3(X1,X2,X3),X2)|~r2_hidden(esk4_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_16, c_0_11])).
cnf(c_0_25, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=X1|r2_hidden(esk4_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_27, plain, ![X15, X16, X17, X18, X19, X20, X21, X22]:(((~r2_hidden(X18,X17)|(X18=X15|X18=X16)|X17!=k2_tarski(X15,X16))&((X19!=X15|r2_hidden(X19,X17)|X17!=k2_tarski(X15,X16))&(X19!=X16|r2_hidden(X19,X17)|X17!=k2_tarski(X15,X16))))&(((esk3_3(X20,X21,X22)!=X20|~r2_hidden(esk3_3(X20,X21,X22),X22)|X22=k2_tarski(X20,X21))&(esk3_3(X20,X21,X22)!=X21|~r2_hidden(esk3_3(X20,X21,X22),X22)|X22=k2_tarski(X20,X21)))&(r2_hidden(esk3_3(X20,X21,X22),X22)|(esk3_3(X20,X21,X22)=X20|esk3_3(X20,X21,X22)=X21)|X22=k2_tarski(X20,X21)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_28, plain, (r2_hidden(X1,X4)|X4!=k4_xboole_0(X2,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_18, c_0_11])).
cnf(c_0_29, negated_conjecture, (k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)))!=k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_20]), c_0_20]), c_0_21]), c_0_21]), c_0_22]), c_0_22]), c_0_11])).
cnf(c_0_30, plain, (X3=k4_xboole_0(X1,k4_xboole_0(X1,X2))|r2_hidden(esk4_3(X1,X2,X3),X3)|r2_hidden(esk4_3(X1,X2,X3),X2)), inference(rw,[status(thm)],[c_0_23, c_0_11])).
cnf(c_0_31, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=X1|~r2_hidden(esk4_3(X1,X2,X1),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_25])).
cnf(c_0_32, plain, (r2_hidden(X1,X2)|X3!=k4_xboole_0(X4,k4_xboole_0(X4,X2))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_26, c_0_11])).
cnf(c_0_33, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_34, plain, (r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_28])).
cnf(c_0_35, negated_conjecture, (r2_hidden(esk4_3(X1,X2,k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)))),k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))))|r2_hidden(esk4_3(X1,X2,k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)))),X2)|k4_xboole_0(X1,k4_xboole_0(X1,X2))!=k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_36, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X1))=X1), inference(spm,[status(thm)],[c_0_31, c_0_25])).
cnf(c_0_37, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X3,k4_xboole_0(X3,X2)))), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_38, plain, (X1=X4|X1=X3|X2!=k2_enumset1(X3,X3,X3,X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_21]), c_0_22])).
cnf(c_0_39, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X3,k4_xboole_0(X3,X4))|~r2_hidden(esk4_3(X3,X4,k4_xboole_0(X1,k4_xboole_0(X1,X2))),X4)|~r2_hidden(esk4_3(X3,X4,k4_xboole_0(X1,k4_xboole_0(X1,X2))),X3)|~r2_hidden(esk4_3(X3,X4,k4_xboole_0(X1,k4_xboole_0(X1,X2))),X2)|~r2_hidden(esk4_3(X3,X4,k4_xboole_0(X1,k4_xboole_0(X1,X2))),X1)), inference(spm,[status(thm)],[c_0_24, c_0_34])).
cnf(c_0_40, negated_conjecture, (r2_hidden(esk4_3(k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)))),k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36])]), c_0_37])).
cnf(c_0_41, plain, (X1=X2|X1=X3|~r2_hidden(X1,k2_enumset1(X3,X3,X3,X2))), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_42, negated_conjecture, (~r2_hidden(esk4_3(k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)))),esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_36]), c_0_40])]), c_0_29])).
cnf(c_0_43, negated_conjecture, (esk4_3(k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))))=esk1_0), inference(spm,[status(thm)],[c_0_41, c_0_40])).
cnf(c_0_44, negated_conjecture, (r2_hidden(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_45, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_43]), c_0_44])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 46
# Proof object clause steps            : 31
# Proof object formula steps           : 15
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 27
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 14
# Processed clauses                    : 463
# ...of these trivial                  : 0
# ...subsumed                          : 200
# ...remaining for further processing  : 263
# Other redundant clauses eliminated   : 437
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 35
# Backward-rewritten                   : 28
# Generated clauses                    : 14610
# ...of the previous two non-trivial   : 13915
# Contextual simplify-reflections      : 17
# Paramodulations                      : 14097
# Factorizations                       : 76
# Equation resolutions                 : 437
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 178
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 171
# Current number of unprocessed clauses: 13466
# ...number of literals in the above   : 80888
# Current number of archived formulas  : 0
# Current number of archived clauses   : 83
# Clause-clause subsumption calls (NU) : 13739
# Rec. Clause-clause subsumption calls : 3674
# Non-unit clause-clause subsumptions  : 240
# Unit Clause-clause subsumption calls : 135
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 11
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 471182

# -------------------------------------------------
# User time                : 0.260 s
# System time              : 0.010 s
# Total time               : 0.270 s
# Maximum resident set size: 28648 pages
