# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t48_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:((~(v1_xboole_0(X2))&(v1_waybel_0(X2,X1)&(v12_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v14_waybel_0(X2,X1)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&X2=k5_waybel_0(X1,X3))))), file('data/MPTP/proofs/2efae9ddc98b4c55bf80d867779b45d1.E_input', t48_waybel_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/2efae9ddc98b4c55bf80d867779b45d1.E_input', t24_orders_2)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/2efae9ddc98b4c55bf80d867779b45d1.E_input', l3_subset_1)).
fof(dt_k5_waybel_0, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/2efae9ddc98b4c55bf80d867779b45d1.E_input', dt_k5_waybel_0)).
fof(t17_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2))))), file('data/MPTP/proofs/2efae9ddc98b4c55bf80d867779b45d1.E_input', t17_waybel_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/2efae9ddc98b4c55bf80d867779b45d1.E_input', t4_subset)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/2efae9ddc98b4c55bf80d867779b45d1.E_input', d3_tarski)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/2efae9ddc98b4c55bf80d867779b45d1.E_input', d10_xboole_0)).
fof(fraenkel_a_2_0_waybel_0, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&l1_orders_2(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r2_hidden(X1,a_2_0_waybel_0(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X4,X5)&r2_hidden(X5,X3))))))), file('data/MPTP/proofs/2efae9ddc98b4c55bf80d867779b45d1.E_input', fraenkel_a_2_0_waybel_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/2efae9ddc98b4c55bf80d867779b45d1.E_input', t3_subset)).
fof(t31_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>r2_lattice3(X1,k3_waybel_0(X1,X2),X3))))), file('data/MPTP/proofs/2efae9ddc98b4c55bf80d867779b45d1.E_input', t31_waybel_0)).
fof(t21_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)=>r1_tarski(k5_waybel_0(X1,X2),k5_waybel_0(X1,X3)))))), file('data/MPTP/proofs/2efae9ddc98b4c55bf80d867779b45d1.E_input', t21_waybel_0)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/2efae9ddc98b4c55bf80d867779b45d1.E_input', d9_lattice3)).
fof(t23_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v12_waybel_0(X2,X1)<=>r1_tarski(k3_waybel_0(X1,X2),X2)))), file('data/MPTP/proofs/2efae9ddc98b4c55bf80d867779b45d1.E_input', t23_waybel_0)).
fof(d15_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k3_waybel_0(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>?[X5]:(m1_subset_1(X5,u1_struct_0(X1))&(r1_orders_2(X1,X4,X5)&r2_hidden(X5,X2))))))))), file('data/MPTP/proofs/2efae9ddc98b4c55bf80d867779b45d1.E_input', d15_waybel_0)).
fof(t14_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k3_waybel_0(X1,X2)=a_2_0_waybel_0(X1,X2))), file('data/MPTP/proofs/2efae9ddc98b4c55bf80d867779b45d1.E_input', t14_waybel_0)).
fof(d21_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:((~(v1_xboole_0(X2))&(v1_waybel_0(X2,X1)&(v12_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v14_waybel_0(X2,X1)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_hidden(X3,X2)&r2_lattice3(X1,X2,X3)))))), file('data/MPTP/proofs/2efae9ddc98b4c55bf80d867779b45d1.E_input', d21_waybel_0)).
fof(c_0_17, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:((~(v1_xboole_0(X2))&(v1_waybel_0(X2,X1)&(v12_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v14_waybel_0(X2,X1)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&X2=k5_waybel_0(X1,X3)))))), inference(assume_negation,[status(cth)],[t48_waybel_0])).
fof(c_0_18, plain, ![X15, X16]:(v2_struct_0(X15)|(~v3_orders_2(X15)|~l1_orders_2(X15))|(~m1_subset_1(X16,u1_struct_0(X15))|r1_orders_2(X15,X16,X16))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_19, negated_conjecture, ![X8]:((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0))))&((~v1_xboole_0(esk2_0)&(v1_waybel_0(esk2_0,esk1_0)&(v12_waybel_0(esk2_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))))&((~v14_waybel_0(esk2_0,esk1_0)|(~m1_subset_1(X8,u1_struct_0(esk1_0))|esk2_0!=k5_waybel_0(esk1_0,X8)))&((m1_subset_1(esk3_0,u1_struct_0(esk1_0))|v14_waybel_0(esk2_0,esk1_0))&(esk2_0=k5_waybel_0(esk1_0,esk3_0)|v14_waybel_0(esk2_0,esk1_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])])])).
fof(c_0_20, plain, ![X65, X66, X67]:(~m1_subset_1(X66,k1_zfmisc_1(X65))|(~r2_hidden(X67,X66)|r2_hidden(X67,X65))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
fof(c_0_21, plain, ![X29, X30]:(v2_struct_0(X29)|~l1_orders_2(X29)|~m1_subset_1(X30,u1_struct_0(X29))|m1_subset_1(k5_waybel_0(X29,X30),k1_zfmisc_1(u1_struct_0(X29)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_waybel_0])])])).
fof(c_0_22, plain, ![X12, X13, X14]:((~r2_hidden(X14,k5_waybel_0(X12,X13))|r1_orders_2(X12,X14,X13)|~m1_subset_1(X14,u1_struct_0(X12))|~m1_subset_1(X13,u1_struct_0(X12))|(v2_struct_0(X12)|~l1_orders_2(X12)))&(~r1_orders_2(X12,X14,X13)|r2_hidden(X14,k5_waybel_0(X12,X13))|~m1_subset_1(X14,u1_struct_0(X12))|~m1_subset_1(X13,u1_struct_0(X12))|(v2_struct_0(X12)|~l1_orders_2(X12)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_waybel_0])])])])])).
cnf(c_0_23, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (v2_struct_0(X1)|m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, plain, (r2_hidden(X2,k5_waybel_0(X1,X3))|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])]), c_0_26])).
fof(c_0_31, plain, ![X40, X41, X42]:(~r2_hidden(X40,X41)|~m1_subset_1(X41,k1_zfmisc_1(X42))|m1_subset_1(X40,X42)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_32, plain, ![X57, X58, X59, X60, X61]:((~r1_tarski(X57,X58)|(~r2_hidden(X59,X57)|r2_hidden(X59,X58)))&((r2_hidden(esk11_2(X60,X61),X60)|r1_tarski(X60,X61))&(~r2_hidden(esk11_2(X60,X61),X61)|r1_tarski(X60,X61)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_33, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k5_waybel_0(X2,X3))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34, negated_conjecture, (r2_hidden(X1,k5_waybel_0(esk1_0,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_25])]), c_0_26])).
cnf(c_0_35, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_36, plain, (r1_tarski(X1,X2)|~r2_hidden(esk11_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_37, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_25])]), c_0_26])).
cnf(c_0_38, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k5_waybel_0(X2,X3))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_35, c_0_28])).
cnf(c_0_39, plain, (r2_hidden(esk11_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_40, plain, ![X10, X11]:(((r1_tarski(X10,X11)|X10!=X11)&(r1_tarski(X11,X10)|X10!=X11))&(~r1_tarski(X10,X11)|~r1_tarski(X11,X10)|X10=X11)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_41, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~m1_subset_1(esk11_2(X1,u1_struct_0(esk1_0)),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_42, plain, (r1_tarski(k5_waybel_0(X1,X2),X3)|m1_subset_1(esk11_2(k5_waybel_0(X1,X2),X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
fof(c_0_43, plain, ![X17, X18, X19, X22, X23]:(((m1_subset_1(esk4_3(X17,X18,X19),u1_struct_0(X18))|~r2_hidden(X17,a_2_0_waybel_0(X18,X19))|(v2_struct_0(X18)|~l1_orders_2(X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))&((X17=esk4_3(X17,X18,X19)|~r2_hidden(X17,a_2_0_waybel_0(X18,X19))|(v2_struct_0(X18)|~l1_orders_2(X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))&((m1_subset_1(esk5_3(X17,X18,X19),u1_struct_0(X18))|~r2_hidden(X17,a_2_0_waybel_0(X18,X19))|(v2_struct_0(X18)|~l1_orders_2(X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))&((r1_orders_2(X18,esk4_3(X17,X18,X19),esk5_3(X17,X18,X19))|~r2_hidden(X17,a_2_0_waybel_0(X18,X19))|(v2_struct_0(X18)|~l1_orders_2(X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))&(r2_hidden(esk5_3(X17,X18,X19),X19)|~r2_hidden(X17,a_2_0_waybel_0(X18,X19))|(v2_struct_0(X18)|~l1_orders_2(X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))))))&(~m1_subset_1(X22,u1_struct_0(X18))|(X17!=X22|(~m1_subset_1(X23,u1_struct_0(X18))|(~r1_orders_2(X18,X22,X23)|~r2_hidden(X23,X19))))|r2_hidden(X17,a_2_0_waybel_0(X18,X19))|(v2_struct_0(X18)|~l1_orders_2(X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_0_waybel_0])])])])])])).
fof(c_0_44, plain, ![X63, X64]:((~m1_subset_1(X63,k1_zfmisc_1(X64))|r1_tarski(X63,X64))&(~r1_tarski(X63,X64)|m1_subset_1(X63,k1_zfmisc_1(X64)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_45, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_46, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_47, negated_conjecture, (r1_tarski(k5_waybel_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_25])]), c_0_26])).
cnf(c_0_48, plain, (r1_orders_2(X1,esk4_3(X2,X1,X3),esk5_3(X2,X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~l1_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_50, plain, (m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_51, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_52, plain, (r2_hidden(X3,a_2_0_waybel_0(X2,X5))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~m1_subset_1(X4,u1_struct_0(X2))|~r1_orders_2(X2,X1,X4)|~r2_hidden(X4,X5)|~l1_orders_2(X2)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_53, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_54, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_45])).
cnf(c_0_55, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k5_waybel_0(esk1_0,X2))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_56, plain, (r2_hidden(esk4_3(X1,X2,X3),k5_waybel_0(X2,esk5_3(X1,X2,X3)))|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_48]), c_0_49]), c_0_50])).
cnf(c_0_57, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_35, c_0_51])).
cnf(c_0_58, plain, (r2_hidden(esk5_3(X1,X2,X3),X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_59, plain, ![X54, X55, X56]:((~r2_lattice3(X54,X55,X56)|r2_lattice3(X54,k3_waybel_0(X54,X55),X56)|~m1_subset_1(X56,u1_struct_0(X54))|~m1_subset_1(X55,k1_zfmisc_1(u1_struct_0(X54)))|(v2_struct_0(X54)|(~v3_orders_2(X54)|(~v4_orders_2(X54)|~l1_orders_2(X54)))))&(~r2_lattice3(X54,k3_waybel_0(X54,X55),X56)|r2_lattice3(X54,X55,X56)|~m1_subset_1(X56,u1_struct_0(X54))|~m1_subset_1(X55,k1_zfmisc_1(u1_struct_0(X54)))|(v2_struct_0(X54)|(~v3_orders_2(X54)|(~v4_orders_2(X54)|~l1_orders_2(X54)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t31_waybel_0])])])])])).
cnf(c_0_60, plain, (r2_hidden(X1,a_2_0_waybel_0(X2,X3))|v2_struct_0(X2)|~r1_orders_2(X2,X1,X4)|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_52, c_0_35])])).
cnf(c_0_61, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_62, negated_conjecture, (r2_hidden(esk4_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|~r2_hidden(X1,a_2_0_waybel_0(esk1_0,X2))|~m1_subset_1(esk5_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_25])]), c_0_26])).
cnf(c_0_63, negated_conjecture, (m1_subset_1(esk5_3(X1,X2,esk2_0),u1_struct_0(esk1_0))|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
fof(c_0_64, plain, ![X31, X32, X33]:(v2_struct_0(X31)|(~v4_orders_2(X31)|~l1_orders_2(X31))|(~m1_subset_1(X32,u1_struct_0(X31))|(~m1_subset_1(X33,u1_struct_0(X31))|(~r1_orders_2(X31,X32,X33)|r1_tarski(k5_waybel_0(X31,X32),k5_waybel_0(X31,X33)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t21_waybel_0])])])])).
fof(c_0_65, plain, ![X24, X25, X26, X27]:((~r2_lattice3(X24,X25,X26)|(~m1_subset_1(X27,u1_struct_0(X24))|(~r2_hidden(X27,X25)|r1_orders_2(X24,X27,X26)))|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))&((m1_subset_1(esk6_3(X24,X25,X26),u1_struct_0(X24))|r2_lattice3(X24,X25,X26)|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))&((r2_hidden(esk6_3(X24,X25,X26),X25)|r2_lattice3(X24,X25,X26)|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))&(~r1_orders_2(X24,esk6_3(X24,X25,X26),X26)|r2_lattice3(X24,X25,X26)|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
cnf(c_0_66, plain, (r2_lattice3(X1,k3_waybel_0(X1,X2),X3)|v2_struct_0(X1)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_67, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_68, plain, ![X52, X53]:((~v12_waybel_0(X53,X52)|r1_tarski(k3_waybel_0(X52,X53),X53)|~m1_subset_1(X53,k1_zfmisc_1(u1_struct_0(X52)))|~l1_orders_2(X52))&(~r1_tarski(k3_waybel_0(X52,X53),X53)|v12_waybel_0(X53,X52)|~m1_subset_1(X53,k1_zfmisc_1(u1_struct_0(X52)))|~l1_orders_2(X52))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_waybel_0])])])])).
cnf(c_0_69, negated_conjecture, (r2_hidden(X1,a_2_0_waybel_0(esk1_0,X2))|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_30]), c_0_25])]), c_0_26]), c_0_35])).
cnf(c_0_70, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_35, c_0_61])).
cnf(c_0_71, negated_conjecture, (r2_hidden(esk4_3(X1,esk1_0,esk2_0),u1_struct_0(esk1_0))|~r2_hidden(X1,a_2_0_waybel_0(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_51]), c_0_25])]), c_0_26])).
cnf(c_0_72, plain, (v2_struct_0(X1)|r1_tarski(k5_waybel_0(X1,X2),k5_waybel_0(X1,X3))|~v4_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_73, plain, (r1_orders_2(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_74, negated_conjecture, (r2_lattice3(esk1_0,k3_waybel_0(esk1_0,X1),X2)|~r2_lattice3(esk1_0,X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_25]), c_0_24])]), c_0_26])).
fof(c_0_75, plain, ![X43, X44, X45, X46, X48, X50]:((((m1_subset_1(esk8_4(X43,X44,X45,X46),u1_struct_0(X43))|~r2_hidden(X46,X45)|~m1_subset_1(X46,u1_struct_0(X43))|X45!=k3_waybel_0(X43,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X43)))|~m1_subset_1(X44,k1_zfmisc_1(u1_struct_0(X43)))|~l1_orders_2(X43))&((r1_orders_2(X43,X46,esk8_4(X43,X44,X45,X46))|~r2_hidden(X46,X45)|~m1_subset_1(X46,u1_struct_0(X43))|X45!=k3_waybel_0(X43,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X43)))|~m1_subset_1(X44,k1_zfmisc_1(u1_struct_0(X43)))|~l1_orders_2(X43))&(r2_hidden(esk8_4(X43,X44,X45,X46),X44)|~r2_hidden(X46,X45)|~m1_subset_1(X46,u1_struct_0(X43))|X45!=k3_waybel_0(X43,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X43)))|~m1_subset_1(X44,k1_zfmisc_1(u1_struct_0(X43)))|~l1_orders_2(X43))))&(~m1_subset_1(X48,u1_struct_0(X43))|(~r1_orders_2(X43,X46,X48)|~r2_hidden(X48,X44))|r2_hidden(X46,X45)|~m1_subset_1(X46,u1_struct_0(X43))|X45!=k3_waybel_0(X43,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X43)))|~m1_subset_1(X44,k1_zfmisc_1(u1_struct_0(X43)))|~l1_orders_2(X43)))&((m1_subset_1(esk9_3(X43,X44,X45),u1_struct_0(X43))|X45=k3_waybel_0(X43,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X43)))|~m1_subset_1(X44,k1_zfmisc_1(u1_struct_0(X43)))|~l1_orders_2(X43))&((~r2_hidden(esk9_3(X43,X44,X45),X45)|(~m1_subset_1(X50,u1_struct_0(X43))|(~r1_orders_2(X43,esk9_3(X43,X44,X45),X50)|~r2_hidden(X50,X44)))|X45=k3_waybel_0(X43,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X43)))|~m1_subset_1(X44,k1_zfmisc_1(u1_struct_0(X43)))|~l1_orders_2(X43))&((m1_subset_1(esk10_3(X43,X44,X45),u1_struct_0(X43))|r2_hidden(esk9_3(X43,X44,X45),X45)|X45=k3_waybel_0(X43,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X43)))|~m1_subset_1(X44,k1_zfmisc_1(u1_struct_0(X43)))|~l1_orders_2(X43))&((r1_orders_2(X43,esk9_3(X43,X44,X45),esk10_3(X43,X44,X45))|r2_hidden(esk9_3(X43,X44,X45),X45)|X45=k3_waybel_0(X43,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X43)))|~m1_subset_1(X44,k1_zfmisc_1(u1_struct_0(X43)))|~l1_orders_2(X43))&(r2_hidden(esk10_3(X43,X44,X45),X44)|r2_hidden(esk9_3(X43,X44,X45),X45)|X45=k3_waybel_0(X43,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X43)))|~m1_subset_1(X44,k1_zfmisc_1(u1_struct_0(X43)))|~l1_orders_2(X43))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d15_waybel_0])])])])])).
cnf(c_0_76, plain, (r1_tarski(k3_waybel_0(X2,X1),X1)|~v12_waybel_0(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_77, negated_conjecture, (v12_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_78, negated_conjecture, (r1_tarski(X1,a_2_0_waybel_0(esk1_0,X2))|~r2_hidden(esk11_2(X1,a_2_0_waybel_0(esk1_0,X2)),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_36, c_0_69])).
fof(c_0_79, plain, ![X34, X35]:(v2_struct_0(X34)|~l1_orders_2(X34)|(~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X34)))|k3_waybel_0(X34,X35)=a_2_0_waybel_0(X34,X35))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_waybel_0])])])])).
cnf(c_0_80, negated_conjecture, (m1_subset_1(esk4_3(X1,esk1_0,esk2_0),u1_struct_0(esk1_0))|~r2_hidden(X1,a_2_0_waybel_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_81, plain, (X1=esk4_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_82, plain, (r1_tarski(k5_waybel_0(X1,esk4_3(X2,X1,X3)),k5_waybel_0(X1,esk5_3(X2,X1,X3)))|v2_struct_0(X1)|~r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v4_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_48]), c_0_49]), c_0_50])).
cnf(c_0_83, negated_conjecture, (r1_orders_2(esk1_0,X1,X2)|~r2_lattice3(esk1_0,X3,X2)|~r2_hidden(X1,k3_waybel_0(esk1_0,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_25])])).
fof(c_0_84, plain, ![X36, X37, X39]:(((m1_subset_1(esk7_2(X36,X37),u1_struct_0(X36))|~v14_waybel_0(X37,X36)|(v1_xboole_0(X37)|(~v1_waybel_0(X37,X36)|(~v12_waybel_0(X37,X36)|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X36))))))|(v2_struct_0(X36)|(~v3_orders_2(X36)|(~v4_orders_2(X36)|~l1_orders_2(X36)))))&((r2_hidden(esk7_2(X36,X37),X37)|~v14_waybel_0(X37,X36)|(v1_xboole_0(X37)|(~v1_waybel_0(X37,X36)|(~v12_waybel_0(X37,X36)|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X36))))))|(v2_struct_0(X36)|(~v3_orders_2(X36)|(~v4_orders_2(X36)|~l1_orders_2(X36)))))&(r2_lattice3(X36,X37,esk7_2(X36,X37))|~v14_waybel_0(X37,X36)|(v1_xboole_0(X37)|(~v1_waybel_0(X37,X36)|(~v12_waybel_0(X37,X36)|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X36))))))|(v2_struct_0(X36)|(~v3_orders_2(X36)|(~v4_orders_2(X36)|~l1_orders_2(X36)))))))&(~m1_subset_1(X39,u1_struct_0(X36))|(~r2_hidden(X39,X37)|~r2_lattice3(X36,X37,X39))|v14_waybel_0(X37,X36)|(v1_xboole_0(X37)|(~v1_waybel_0(X37,X36)|(~v12_waybel_0(X37,X36)|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X36))))))|(v2_struct_0(X36)|(~v3_orders_2(X36)|(~v4_orders_2(X36)|~l1_orders_2(X36)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d21_waybel_0])])])])])])).
cnf(c_0_85, plain, (r1_orders_2(X2,X1,X3)|v2_struct_0(X2)|~r2_hidden(X1,k5_waybel_0(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_86, negated_conjecture, (esk2_0=k5_waybel_0(esk1_0,esk3_0)|v14_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_87, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))|v14_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_88, plain, (r1_orders_2(X1,esk9_3(X1,X2,X3),esk10_3(X1,X2,X3))|r2_hidden(esk9_3(X1,X2,X3),X3)|X3=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_89, plain, (m1_subset_1(esk9_3(X1,X2,X3),u1_struct_0(X1))|X3=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_90, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_91, negated_conjecture, (r1_tarski(k3_waybel_0(esk1_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_51]), c_0_25])])).
cnf(c_0_92, negated_conjecture, (r1_tarski(X1,a_2_0_waybel_0(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_78, c_0_39])).
cnf(c_0_93, plain, (v2_struct_0(X1)|k3_waybel_0(X1,X2)=a_2_0_waybel_0(X1,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_94, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,a_2_0_waybel_0(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_51]), c_0_25])]), c_0_26])).
cnf(c_0_95, plain, (r2_hidden(X1,k5_waybel_0(X2,esk5_3(X3,X2,X4)))|v2_struct_0(X2)|~r2_hidden(X1,k5_waybel_0(X2,esk4_3(X3,X2,X4)))|~r2_hidden(X3,a_2_0_waybel_0(X2,X4))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v4_orders_2(X2)), inference(spm,[status(thm)],[c_0_46, c_0_82])).
cnf(c_0_96, negated_conjecture, (r1_orders_2(esk1_0,esk11_2(k3_waybel_0(esk1_0,X1),X2),X3)|r1_tarski(k3_waybel_0(esk1_0,X1),X2)|~r2_lattice3(esk1_0,X1,X3)|~m1_subset_1(esk11_2(k3_waybel_0(esk1_0,X1),X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_83, c_0_39])).
cnf(c_0_97, plain, (r2_lattice3(X1,X2,esk7_2(X1,X2))|v1_xboole_0(X2)|v2_struct_0(X1)|~v14_waybel_0(X2,X1)|~v1_waybel_0(X2,X1)|~v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_98, plain, (r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk6_3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_99, negated_conjecture, (r1_orders_2(esk1_0,X1,esk3_0)|v14_waybel_0(esk2_0,esk1_0)|~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_25])]), c_0_26]), c_0_57]), c_0_87])).
cnf(c_0_100, plain, (r2_hidden(X3,X5)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,X3,X1)|~r2_hidden(X1,X4)|~m1_subset_1(X3,u1_struct_0(X2))|X5!=k3_waybel_0(X2,X4)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_101, negated_conjecture, (r2_hidden(esk4_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|~r2_hidden(X1,a_2_0_waybel_0(esk1_0,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_50]), c_0_25])]), c_0_26])).
cnf(c_0_102, plain, (X1=k3_waybel_0(X2,X3)|r2_hidden(esk9_3(X2,X3,X1),a_2_0_waybel_0(X2,X4))|r2_hidden(esk9_3(X2,X3,X1),X1)|v2_struct_0(X2)|~r2_hidden(esk10_3(X2,X3,X1),X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_88]), c_0_89])).
cnf(c_0_103, plain, (r2_hidden(esk10_3(X1,X2,X3),X2)|r2_hidden(esk9_3(X1,X2,X3),X3)|X3=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_104, negated_conjecture, (k3_waybel_0(esk1_0,esk2_0)=esk2_0|~r1_tarski(esk2_0,k3_waybel_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_105, negated_conjecture, (r1_tarski(X1,k3_waybel_0(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_25])]), c_0_26])).
cnf(c_0_106, negated_conjecture, (r1_tarski(a_2_0_waybel_0(esk1_0,esk2_0),X1)|m1_subset_1(esk11_2(a_2_0_waybel_0(esk1_0,esk2_0),X1),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_94, c_0_39])).
cnf(c_0_107, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k5_waybel_0(esk1_0,esk4_3(X2,esk1_0,X3)))|~r2_hidden(X2,a_2_0_waybel_0(esk1_0,X3))|~m1_subset_1(esk5_3(X2,esk1_0,X3),u1_struct_0(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_95]), c_0_25]), c_0_67])]), c_0_26])).
cnf(c_0_108, negated_conjecture, (r1_orders_2(esk1_0,esk11_2(k3_waybel_0(esk1_0,X1),X2),esk7_2(esk1_0,X1))|r1_tarski(k3_waybel_0(esk1_0,X1),X2)|v1_xboole_0(X1)|~v14_waybel_0(X1,esk1_0)|~m1_subset_1(esk11_2(k3_waybel_0(esk1_0,X1),X2),u1_struct_0(esk1_0))|~m1_subset_1(esk7_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v12_waybel_0(X1,esk1_0)|~v1_waybel_0(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_25]), c_0_67]), c_0_24])]), c_0_26])).
cnf(c_0_109, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,k3_waybel_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_46, c_0_91])).
cnf(c_0_110, plain, (v14_waybel_0(X3,X2)|v1_xboole_0(X3)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r2_lattice3(X2,X3,X1)|~v1_waybel_0(X3,X2)|~v12_waybel_0(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v3_orders_2(X2)|~v4_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_111, negated_conjecture, (r2_lattice3(esk1_0,X1,esk3_0)|v14_waybel_0(esk2_0,esk1_0)|~r2_hidden(esk6_3(esk1_0,X1,esk3_0),esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_25])]), c_0_87])).
cnf(c_0_112, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_113, plain, (r2_hidden(X1,X2)|X2!=k3_waybel_0(X3,X4)|~r1_orders_2(X3,X1,X5)|~r2_hidden(X5,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)), inference(csr,[status(thm)],[c_0_100, c_0_35])).
cnf(c_0_114, negated_conjecture, (m1_subset_1(esk4_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|~r2_hidden(X1,a_2_0_waybel_0(esk1_0,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_70, c_0_101])).
cnf(c_0_115, plain, (X3=k3_waybel_0(X1,X2)|~r2_hidden(esk9_3(X1,X2,X3),X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,esk9_3(X1,X2,X3),X4)|~r2_hidden(X4,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_116, plain, (X1=k3_waybel_0(X2,X3)|r2_hidden(esk9_3(X2,X3,X1),a_2_0_waybel_0(X2,X3))|r2_hidden(esk9_3(X2,X3,X1),X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_117, negated_conjecture, (k3_waybel_0(esk1_0,esk2_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_51])])).
cnf(c_0_118, negated_conjecture, (r1_tarski(a_2_0_waybel_0(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_41, c_0_106])).
cnf(c_0_119, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k5_waybel_0(esk1_0,X2))|~r2_hidden(X2,a_2_0_waybel_0(esk1_0,X3))|~m1_subset_1(esk5_3(X2,esk1_0,X3),u1_struct_0(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_81]), c_0_25])]), c_0_26])).
cnf(c_0_120, negated_conjecture, (r2_hidden(esk11_2(k3_waybel_0(esk1_0,X1),X2),k5_waybel_0(esk1_0,esk7_2(esk1_0,X1)))|r1_tarski(k3_waybel_0(esk1_0,X1),X2)|v1_xboole_0(X1)|~v14_waybel_0(X1,esk1_0)|~m1_subset_1(esk11_2(k3_waybel_0(esk1_0,X1),X2),u1_struct_0(esk1_0))|~m1_subset_1(esk7_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v12_waybel_0(X1,esk1_0)|~v1_waybel_0(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_108]), c_0_25])]), c_0_26])).
cnf(c_0_121, negated_conjecture, (r2_hidden(esk11_2(k3_waybel_0(esk1_0,esk2_0),X1),esk2_0)|r1_tarski(k3_waybel_0(esk1_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_109, c_0_39])).
cnf(c_0_122, plain, (v14_waybel_0(X1,X2)|v1_xboole_0(X1)|v2_struct_0(X2)|~r2_lattice3(X2,X1,X3)|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v12_waybel_0(X1,X2)|~v1_waybel_0(X1,X2)|~l1_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[c_0_110, c_0_35])).
cnf(c_0_123, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk3_0)|v14_waybel_0(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_25])]), c_0_87])).
cnf(c_0_124, negated_conjecture, (v1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_125, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_126, negated_conjecture, (r2_hidden(esk3_0,esk2_0)|v14_waybel_0(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_86]), c_0_87])).
cnf(c_0_127, negated_conjecture, (r2_hidden(X1,esk2_0)|esk2_0!=k3_waybel_0(esk1_0,X2)|~r1_orders_2(esk1_0,X1,X3)|~r2_hidden(X3,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_51]), c_0_25])])).
cnf(c_0_128, plain, (r1_orders_2(X1,X2,esk5_3(X2,X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_48, c_0_81])).
cnf(c_0_129, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,a_2_0_waybel_0(esk1_0,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_81]), c_0_25])]), c_0_26])).
cnf(c_0_130, plain, (X1=k3_waybel_0(X2,X3)|~r1_orders_2(X2,esk9_3(X2,X3,X1),X4)|~r2_hidden(esk9_3(X2,X3,X1),X1)|~r2_hidden(X4,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(csr,[status(thm)],[c_0_115, c_0_35])).
cnf(c_0_131, negated_conjecture, (X1=esk2_0|r2_hidden(esk9_3(esk1_0,esk2_0,X1),a_2_0_waybel_0(esk1_0,esk2_0))|r2_hidden(esk9_3(esk1_0,esk2_0,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_51]), c_0_117]), c_0_25])]), c_0_26])).
cnf(c_0_132, negated_conjecture, (m1_subset_1(a_2_0_waybel_0(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_53, c_0_118])).
cnf(c_0_133, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X2,a_2_0_waybel_0(esk1_0,esk2_0))|~r2_hidden(X1,k5_waybel_0(esk1_0,X2))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_63]), c_0_51]), c_0_25])]), c_0_26])).
cnf(c_0_134, negated_conjecture, (r1_tarski(k3_waybel_0(esk1_0,X1),k5_waybel_0(esk1_0,esk7_2(esk1_0,X1)))|v1_xboole_0(X1)|~v14_waybel_0(X1,esk1_0)|~m1_subset_1(esk11_2(k3_waybel_0(esk1_0,X1),k5_waybel_0(esk1_0,esk7_2(esk1_0,X1))),u1_struct_0(esk1_0))|~m1_subset_1(esk7_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v12_waybel_0(X1,esk1_0)|~v1_waybel_0(X1,esk1_0)), inference(spm,[status(thm)],[c_0_36, c_0_120])).
cnf(c_0_135, negated_conjecture, (r1_tarski(k3_waybel_0(esk1_0,esk2_0),X1)|m1_subset_1(esk11_2(k3_waybel_0(esk1_0,esk2_0),X1),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_57, c_0_121])).
cnf(c_0_136, negated_conjecture, (v14_waybel_0(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_51]), c_0_77]), c_0_124]), c_0_25]), c_0_67]), c_0_24])]), c_0_125]), c_0_26]), c_0_126])).
cnf(c_0_137, negated_conjecture, (r2_hidden(X1,k3_waybel_0(esk1_0,X2))|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_93]), c_0_25])]), c_0_26])).
cnf(c_0_138, plain, (r1_orders_2(X1,esk11_2(k5_waybel_0(X1,X2),X3),X2)|r1_tarski(k5_waybel_0(X1,X2),X3)|v2_struct_0(X1)|~m1_subset_1(esk11_2(k5_waybel_0(X1,X2),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_85, c_0_39])).
cnf(c_0_139, negated_conjecture, (r2_hidden(X1,esk2_0)|esk2_0!=k3_waybel_0(esk1_0,X2)|~r2_hidden(esk5_3(X1,esk1_0,X3),X2)|~r2_hidden(X1,a_2_0_waybel_0(esk1_0,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128]), c_0_25])]), c_0_26]), c_0_129])).
cnf(c_0_140, negated_conjecture, (X1=k3_waybel_0(esk1_0,X2)|~r2_hidden(esk9_3(esk1_0,X2,X1),X1)|~r2_hidden(esk9_3(esk1_0,X2,X1),X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_30]), c_0_25])]), c_0_35])).
cnf(c_0_141, negated_conjecture, (a_2_0_waybel_0(esk1_0,esk2_0)=esk2_0|r2_hidden(esk9_3(esk1_0,esk2_0,a_2_0_waybel_0(esk1_0,esk2_0)),a_2_0_waybel_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_131, c_0_132])).
cnf(c_0_142, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k5_waybel_0(esk1_0,X2))|~r2_hidden(X2,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_93]), c_0_117]), c_0_51]), c_0_25])]), c_0_26])).
cnf(c_0_143, negated_conjecture, (r1_tarski(k3_waybel_0(esk1_0,esk2_0),k5_waybel_0(esk1_0,esk7_2(esk1_0,esk2_0)))|~m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_136]), c_0_51]), c_0_77]), c_0_124])]), c_0_125])).
cnf(c_0_144, negated_conjecture, (~v14_waybel_0(esk2_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|esk2_0!=k5_waybel_0(esk1_0,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_145, negated_conjecture, (r1_tarski(X1,k3_waybel_0(esk1_0,X2))|~r2_hidden(esk11_2(X1,k3_waybel_0(esk1_0,X2)),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_36, c_0_137])).
cnf(c_0_146, plain, (r2_hidden(esk11_2(k5_waybel_0(X1,X2),X3),a_2_0_waybel_0(X1,X4))|r1_tarski(k5_waybel_0(X1,X2),X3)|v2_struct_0(X1)|~r2_hidden(X2,X4)|~m1_subset_1(esk11_2(k5_waybel_0(X1,X2),X3),u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_138]), c_0_35])).
cnf(c_0_147, negated_conjecture, (r2_hidden(X1,esk2_0)|esk2_0!=k3_waybel_0(esk1_0,X2)|~r2_hidden(X1,a_2_0_waybel_0(esk1_0,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_58]), c_0_25])]), c_0_26])).
cnf(c_0_148, negated_conjecture, (a_2_0_waybel_0(esk1_0,esk2_0)=esk2_0|~r2_hidden(esk9_3(esk1_0,esk2_0,a_2_0_waybel_0(esk1_0,esk2_0)),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_141]), c_0_117]), c_0_132]), c_0_51])])).
cnf(c_0_149, negated_conjecture, (r2_hidden(esk11_2(k5_waybel_0(esk1_0,X1),X2),u1_struct_0(esk1_0))|r1_tarski(k5_waybel_0(esk1_0,X1),X2)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_142, c_0_39])).
cnf(c_0_150, negated_conjecture, (k5_waybel_0(esk1_0,esk7_2(esk1_0,esk2_0))=k3_waybel_0(esk1_0,esk2_0)|~r1_tarski(k5_waybel_0(esk1_0,esk7_2(esk1_0,esk2_0)),k3_waybel_0(esk1_0,esk2_0))|~m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_90, c_0_143])).
cnf(c_0_151, negated_conjecture, (k5_waybel_0(esk1_0,X1)!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_144, c_0_136])])).
cnf(c_0_152, negated_conjecture, (r1_tarski(k5_waybel_0(X1,X2),k3_waybel_0(esk1_0,a_2_0_waybel_0(X1,X3)))|v2_struct_0(X1)|~r2_hidden(X2,X3)|~m1_subset_1(esk11_2(k5_waybel_0(X1,X2),k3_waybel_0(esk1_0,a_2_0_waybel_0(X1,X3))),u1_struct_0(X1))|~m1_subset_1(a_2_0_waybel_0(X1,X3),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_145, c_0_146])).
cnf(c_0_153, negated_conjecture, (a_2_0_waybel_0(esk1_0,esk2_0)=esk2_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_141]), c_0_117]), c_0_51])]), c_0_148])).
cnf(c_0_154, negated_conjecture, (r1_tarski(k5_waybel_0(esk1_0,X1),X2)|m1_subset_1(esk11_2(k5_waybel_0(esk1_0,X1),X2),u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_70, c_0_149])).
cnf(c_0_155, negated_conjecture, (~r1_tarski(k5_waybel_0(esk1_0,esk7_2(esk1_0,esk2_0)),esk2_0)|~m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_150, c_0_117]), c_0_117]), c_0_151])).
cnf(c_0_156, negated_conjecture, (r1_tarski(k5_waybel_0(esk1_0,X1),esk2_0)|~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_153]), c_0_117]), c_0_117]), c_0_51]), c_0_25])]), c_0_26]), c_0_154])).
cnf(c_0_157, negated_conjecture, (~r2_hidden(esk7_2(esk1_0,esk2_0),esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_156]), c_0_57])).
cnf(c_0_158, plain, (r2_hidden(esk7_2(X1,X2),X2)|v1_xboole_0(X2)|v2_struct_0(X1)|~v14_waybel_0(X2,X1)|~v1_waybel_0(X2,X1)|~v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_159, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_158]), c_0_136]), c_0_51]), c_0_77]), c_0_124]), c_0_25]), c_0_67]), c_0_24])]), c_0_125]), c_0_26]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 160
# Proof object clause steps            : 125
# Proof object formula steps           : 35
# Proof object conjectures             : 77
# Proof object clause conjectures      : 74
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 44
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 74
# Proof object simplifying inferences  : 161
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 54
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 54
# Processed clauses                    : 13673
# ...of these trivial                  : 698
# ...subsumed                          : 10128
# ...remaining for further processing  : 2847
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 407
# Backward-rewritten                   : 625
# Generated clauses                    : 67833
# ...of the previous two non-trivial   : 65403
# Contextual simplify-reflections      : 725
# Paramodulations                      : 67822
# Factorizations                       : 0
# Equation resolutions                 : 11
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1759
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 14
#    Non-unit-clauses                  : 1728
# Current number of unprocessed clauses: 49911
# ...number of literals in the above   : 464239
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1085
# Clause-clause subsumption calls (NU) : 986587
# Rec. Clause-clause subsumption calls : 92221
# Non-unit clause-clause subsumptions  : 8482
# Unit Clause-clause subsumption calls : 2721
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 57
# BW rewrite match successes           : 17
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2568345

# -------------------------------------------------
# User time                : 1.910 s
# System time              : 0.023 s
# Total time               : 1.933 s
# Maximum resident set size: 59388 pages
