# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t22_waybel_9, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(k9_yellow_6(X1,X2)))=>![X4]:(m1_subset_1(X4,u1_struct_0(k9_yellow_6(X1,X2)))=>m1_subset_1(k3_xboole_0(X3,X4),u1_struct_0(k9_yellow_6(X1,X2))))))), file('data/MPTP/proofs/2f67a156a03f4d9fb8d2860edaa2d8be.E_input', t22_waybel_9)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/2f67a156a03f4d9fb8d2860edaa2d8be.E_input', d4_xboole_0)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/2f67a156a03f4d9fb8d2860edaa2d8be.E_input', t12_setfam_1)).
fof(t4_waybel_7, axiom, ![X1]:u1_struct_0(k3_yellow_1(X1))=k9_setfam_1(X1), file('data/MPTP/proofs/2f67a156a03f4d9fb8d2860edaa2d8be.E_input', t4_waybel_7)).
fof(redefinition_k9_setfam_1, axiom, ![X1]:k9_setfam_1(X1)=k1_zfmisc_1(X1), file('data/MPTP/proofs/2f67a156a03f4d9fb8d2860edaa2d8be.E_input', redefinition_k9_setfam_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/2f67a156a03f4d9fb8d2860edaa2d8be.E_input', d2_subset_1)).
fof(dt_m1_connsp_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,u1_struct_0(X1)))=>![X3]:(m1_connsp_2(X3,X1,X2)=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/2f67a156a03f4d9fb8d2860edaa2d8be.E_input', dt_m1_connsp_2)).
fof(t21_waybel_9, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(k9_yellow_6(X1,X2)))=>m1_connsp_2(X3,X1,X2)))), file('data/MPTP/proofs/2f67a156a03f4d9fb8d2860edaa2d8be.E_input', t21_waybel_9)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/2f67a156a03f4d9fb8d2860edaa2d8be.E_input', d1_xboole_0)).
fof(t39_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X3,u1_struct_0(k9_yellow_6(X1,X2)))<=>(r2_hidden(X2,X3)&v3_pre_topc(X3,X1)))))), file('data/MPTP/proofs/2f67a156a03f4d9fb8d2860edaa2d8be.E_input', t39_yellow_6)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/2f67a156a03f4d9fb8d2860edaa2d8be.E_input', t3_subset)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/2f67a156a03f4d9fb8d2860edaa2d8be.E_input', d3_tarski)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/2f67a156a03f4d9fb8d2860edaa2d8be.E_input', t4_subset)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/2f67a156a03f4d9fb8d2860edaa2d8be.E_input', t1_subset)).
fof(fc6_tops_1, axiom, ![X1, X2, X3]:(((v2_pre_topc(X1)&l1_pre_topc(X1))&((v3_pre_topc(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(v3_pre_topc(X3,X1)&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))))=>v3_pre_topc(k3_xboole_0(X2,X3),X1)), file('data/MPTP/proofs/2f67a156a03f4d9fb8d2860edaa2d8be.E_input', fc6_tops_1)).
fof(c_0_15, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(k9_yellow_6(X1,X2)))=>![X4]:(m1_subset_1(X4,u1_struct_0(k9_yellow_6(X1,X2)))=>m1_subset_1(k3_xboole_0(X3,X4),u1_struct_0(k9_yellow_6(X1,X2)))))))), inference(assume_negation,[status(cth)],[t22_waybel_9])).
fof(c_0_16, plain, ![X42, X43, X44, X45, X46, X47, X48, X49]:((((r2_hidden(X45,X42)|~r2_hidden(X45,X44)|X44!=k3_xboole_0(X42,X43))&(r2_hidden(X45,X43)|~r2_hidden(X45,X44)|X44!=k3_xboole_0(X42,X43)))&(~r2_hidden(X46,X42)|~r2_hidden(X46,X43)|r2_hidden(X46,X44)|X44!=k3_xboole_0(X42,X43)))&((~r2_hidden(esk7_3(X47,X48,X49),X49)|(~r2_hidden(esk7_3(X47,X48,X49),X47)|~r2_hidden(esk7_3(X47,X48,X49),X48))|X49=k3_xboole_0(X47,X48))&((r2_hidden(esk7_3(X47,X48,X49),X47)|r2_hidden(esk7_3(X47,X48,X49),X49)|X49=k3_xboole_0(X47,X48))&(r2_hidden(esk7_3(X47,X48,X49),X48)|r2_hidden(esk7_3(X47,X48,X49),X49)|X49=k3_xboole_0(X47,X48))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_17, plain, ![X40, X41]:k1_setfam_1(k2_tarski(X40,X41))=k3_xboole_0(X40,X41), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_18, plain, ![X54]:u1_struct_0(k3_yellow_1(X54))=k9_setfam_1(X54), inference(variable_rename,[status(thm)],[t4_waybel_7])).
fof(c_0_19, plain, ![X53]:k9_setfam_1(X53)=k1_zfmisc_1(X53), inference(variable_rename,[status(thm)],[redefinition_k9_setfam_1])).
fof(c_0_20, plain, ![X9, X10]:(((~m1_subset_1(X10,X9)|r2_hidden(X10,X9)|v1_xboole_0(X9))&(~r2_hidden(X10,X9)|m1_subset_1(X10,X9)|v1_xboole_0(X9)))&((~m1_subset_1(X10,X9)|v1_xboole_0(X10)|~v1_xboole_0(X9))&(~v1_xboole_0(X10)|m1_subset_1(X10,X9)|~v1_xboole_0(X9)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_21, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))&(m1_subset_1(esk4_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))&~m1_subset_1(k3_xboole_0(esk3_0,esk4_0),u1_struct_0(k9_yellow_6(esk1_0,esk2_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_15])])])])).
cnf(c_0_22, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_24, plain, ![X15, X16, X17]:(v2_struct_0(X15)|(~v2_pre_topc(X15)|~l1_pre_topc(X15))|~m1_subset_1(X16,u1_struct_0(X15))|(~m1_connsp_2(X17,X15,X16)|m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X15))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_connsp_2])])])])).
cnf(c_0_25, plain, (u1_struct_0(k3_yellow_1(X1))=k9_setfam_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, plain, (k9_setfam_1(X1)=k1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_27, plain, ![X37, X38, X39]:(v2_struct_0(X37)|(~v2_pre_topc(X37)|~l1_pre_topc(X37))|(~m1_subset_1(X38,u1_struct_0(X37))|(~m1_subset_1(X39,u1_struct_0(k9_yellow_6(X37,X38)))|m1_connsp_2(X39,X37,X38)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t21_waybel_9])])])])).
cnf(c_0_28, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, plain, (r2_hidden(X1,X2)|X3!=k1_setfam_1(k2_tarski(X4,X2))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_22, c_0_23])).
fof(c_0_31, plain, ![X11, X12, X13]:((~v1_xboole_0(X11)|~r2_hidden(X12,X11))&(r2_hidden(esk5_1(X13),X13)|v1_xboole_0(X13))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_32, plain, ![X23, X24, X25]:(((r2_hidden(X24,X25)|~r2_hidden(X25,u1_struct_0(k9_yellow_6(X23,X24)))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X23)))|~m1_subset_1(X24,u1_struct_0(X23))|(v2_struct_0(X23)|(~v2_pre_topc(X23)|~l1_pre_topc(X23))))&(v3_pre_topc(X25,X23)|~r2_hidden(X25,u1_struct_0(k9_yellow_6(X23,X24)))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X23)))|~m1_subset_1(X24,u1_struct_0(X23))|(v2_struct_0(X23)|(~v2_pre_topc(X23)|~l1_pre_topc(X23)))))&(~r2_hidden(X24,X25)|~v3_pre_topc(X25,X23)|r2_hidden(X25,u1_struct_0(k9_yellow_6(X23,X24)))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X23)))|~m1_subset_1(X24,u1_struct_0(X23))|(v2_struct_0(X23)|(~v2_pre_topc(X23)|~l1_pre_topc(X23))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t39_yellow_6])])])])])).
cnf(c_0_33, plain, (v2_struct_0(X1)|m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_connsp_2(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, plain, (u1_struct_0(k3_yellow_1(X1))=k1_zfmisc_1(X1)), inference(rw,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_35, plain, (v2_struct_0(X1)|m1_connsp_2(X3,X1,X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(k9_yellow_6(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_37, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_38, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_39, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_40, negated_conjecture, (~m1_subset_1(k3_xboole_0(esk3_0,esk4_0),u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_41, plain, (m1_subset_1(X1,X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_42, negated_conjecture, (r2_hidden(esk4_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))|v1_xboole_0(u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_43, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_setfam_1(k2_tarski(X3,X2)))), inference(er,[status(thm)],[c_0_30])).
cnf(c_0_44, plain, (r2_hidden(esk5_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_45, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_46, plain, ![X18, X19]:((~m1_subset_1(X18,k1_zfmisc_1(X19))|r1_tarski(X18,X19))&(~r1_tarski(X18,X19)|m1_subset_1(X18,k1_zfmisc_1(X19)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_47, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_48, plain, (r2_hidden(X1,X2)|v2_struct_0(X3)|~r2_hidden(X2,u1_struct_0(k9_yellow_6(X3,X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_49, plain, (v2_struct_0(X1)|m1_subset_1(X3,u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_connsp_2(X3,X1,X2)|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_50, negated_conjecture, (m1_connsp_2(esk4_0,esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_29]), c_0_36]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_51, negated_conjecture, (~m1_subset_1(k1_setfam_1(k2_tarski(esk3_0,esk4_0)),u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(rw,[status(thm)],[c_0_40, c_0_23])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk4_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))|m1_subset_1(X1,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_53, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_54, plain, (r2_hidden(esk5_1(k1_setfam_1(k2_tarski(X1,X2))),X2)|v1_xboole_0(k1_setfam_1(k2_tarski(X1,X2)))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_55, negated_conjecture, (v1_xboole_0(esk4_0)|~v1_xboole_0(u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_45, c_0_29])).
cnf(c_0_56, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_57, negated_conjecture, (m1_connsp_2(esk3_0,esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_47]), c_0_36]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_58, plain, (v2_struct_0(X3)|r2_hidden(X1,X2)|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~m1_subset_1(X1,u1_struct_0(X3))|~r2_hidden(X2,u1_struct_0(k9_yellow_6(X3,X1)))|~m1_subset_1(X2,u1_struct_0(k3_yellow_1(u1_struct_0(X3))))), inference(rw,[status(thm)],[c_0_48, c_0_34])).
cnf(c_0_59, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_36]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_60, negated_conjecture, (r2_hidden(esk4_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))|~v1_xboole_0(k1_setfam_1(k2_tarski(esk3_0,esk4_0)))), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_61, plain, (v1_xboole_0(k1_setfam_1(k2_tarski(X1,X2)))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_62, negated_conjecture, (r2_hidden(esk4_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))|v1_xboole_0(esk4_0)), inference(spm,[status(thm)],[c_0_55, c_0_42])).
fof(c_0_63, plain, ![X29, X30, X31, X32, X33]:((~r1_tarski(X29,X30)|(~r2_hidden(X31,X29)|r2_hidden(X31,X30)))&((r2_hidden(esk6_2(X32,X33),X32)|r1_tarski(X32,X33))&(~r2_hidden(esk6_2(X32,X33),X33)|r1_tarski(X32,X33)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_64, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,u1_struct_0(k3_yellow_1(X2)))), inference(rw,[status(thm)],[c_0_56, c_0_34])).
cnf(c_0_65, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_57]), c_0_36]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_66, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_67, negated_conjecture, (r2_hidden(X1,esk4_0)|~r2_hidden(esk4_0,u1_struct_0(k9_yellow_6(esk1_0,X1)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_68, negated_conjecture, (r2_hidden(esk4_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62])).
fof(c_0_69, plain, ![X20, X21, X22]:(~r2_hidden(X20,X21)|~m1_subset_1(X21,k1_zfmisc_1(X22))|m1_subset_1(X20,X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_70, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_71, negated_conjecture, (r1_tarski(esk3_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_72, plain, (r2_hidden(X1,X2)|X3!=k1_setfam_1(k2_tarski(X2,X4))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_66, c_0_23])).
cnf(c_0_73, negated_conjecture, (r2_hidden(esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_36])])).
cnf(c_0_74, plain, (r2_hidden(X2,u1_struct_0(k9_yellow_6(X3,X1)))|v2_struct_0(X3)|~r2_hidden(X1,X2)|~v3_pre_topc(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_75, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_76, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_77, plain, (r1_tarski(X1,X2)|~r2_hidden(esk6_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_78, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_79, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_setfam_1(k2_tarski(X2,X3)))), inference(er,[status(thm)],[c_0_72])).
cnf(c_0_80, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_81, negated_conjecture, (v1_xboole_0(esk4_0)|~v1_xboole_0(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0))))), inference(spm,[status(thm)],[c_0_45, c_0_59])).
cnf(c_0_82, negated_conjecture, (r2_hidden(esk4_0,u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0))))|v1_xboole_0(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0))))), inference(spm,[status(thm)],[c_0_28, c_0_59])).
cnf(c_0_83, negated_conjecture, (~v1_xboole_0(esk4_0)), inference(spm,[status(thm)],[c_0_53, c_0_73])).
cnf(c_0_84, plain, (v3_pre_topc(X1,X2)|v2_struct_0(X2)|~r2_hidden(X1,u1_struct_0(k9_yellow_6(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,u1_struct_0(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_85, plain, ![X51, X52]:(~r2_hidden(X51,X52)|m1_subset_1(X51,X52)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_86, plain, (v2_struct_0(X3)|r2_hidden(X2,u1_struct_0(k9_yellow_6(X3,X1)))|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~r2_hidden(X1,X2)|~v3_pre_topc(X2,X3)|~m1_subset_1(X1,u1_struct_0(X3))|~m1_subset_1(X2,u1_struct_0(k3_yellow_1(u1_struct_0(X3))))), inference(rw,[status(thm)],[c_0_74, c_0_34])).
cnf(c_0_87, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,u1_struct_0(k3_yellow_1(X3)))), inference(rw,[status(thm)],[c_0_75, c_0_34])).
cnf(c_0_88, plain, (m1_subset_1(X1,u1_struct_0(k3_yellow_1(X2)))|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_76, c_0_34])).
cnf(c_0_89, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~r2_hidden(esk6_2(X1,u1_struct_0(esk1_0)),esk3_0)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_90, plain, (r1_tarski(k1_setfam_1(k2_tarski(X1,X2)),X3)|r2_hidden(esk6_2(k1_setfam_1(k2_tarski(X1,X2)),X3),X1)), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_91, negated_conjecture, (r2_hidden(esk4_0,u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_83])).
cnf(c_0_92, plain, (v2_struct_0(X2)|v3_pre_topc(X1,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,u1_struct_0(k9_yellow_6(X2,X3)))|~m1_subset_1(X1,u1_struct_0(k3_yellow_1(u1_struct_0(X2))))), inference(rw,[status(thm)],[c_0_84, c_0_34])).
cnf(c_0_93, negated_conjecture, (r2_hidden(esk3_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))|v1_xboole_0(u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_28, c_0_47])).
cnf(c_0_94, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_95, plain, (r2_hidden(X1,u1_struct_0(k9_yellow_6(X2,X3)))|v2_struct_0(X2)|~v3_pre_topc(X1,X2)|~r2_hidden(X3,X1)|~m1_subset_1(X1,u1_struct_0(k3_yellow_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_96, plain, (r2_hidden(X1,u1_struct_0(k3_yellow_1(X2)))|v1_xboole_0(u1_struct_0(k3_yellow_1(X2)))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_28, c_0_88])).
cnf(c_0_97, negated_conjecture, (r1_tarski(k1_setfam_1(k2_tarski(esk3_0,X1)),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_98, negated_conjecture, (~v1_xboole_0(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0))))), inference(spm,[status(thm)],[c_0_53, c_0_91])).
fof(c_0_99, plain, ![X26, X27, X28]:(~v2_pre_topc(X26)|~l1_pre_topc(X26)|(~v3_pre_topc(X27,X26)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X26)))|(~v3_pre_topc(X28,X26)|~m1_subset_1(X28,k1_zfmisc_1(u1_struct_0(X26)))))|v3_pre_topc(k3_xboole_0(X27,X28),X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc6_tops_1])])).
cnf(c_0_100, negated_conjecture, (v3_pre_topc(esk3_0,esk1_0)|~r2_hidden(esk3_0,u1_struct_0(k9_yellow_6(esk1_0,X1)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_65]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_101, negated_conjecture, (r2_hidden(esk3_0,u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))|v1_xboole_0(esk4_0)), inference(spm,[status(thm)],[c_0_55, c_0_93])).
cnf(c_0_102, negated_conjecture, (~r2_hidden(k1_setfam_1(k2_tarski(esk3_0,esk4_0)),u1_struct_0(k9_yellow_6(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_51, c_0_94])).
cnf(c_0_103, plain, (r2_hidden(X1,u1_struct_0(k9_yellow_6(X2,X3)))|v2_struct_0(X2)|~v3_pre_topc(X1,X2)|~r2_hidden(X1,u1_struct_0(k3_yellow_1(u1_struct_0(X2))))|~r2_hidden(X3,X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_95, c_0_94])).
cnf(c_0_104, negated_conjecture, (r2_hidden(k1_setfam_1(k2_tarski(esk3_0,X1)),u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_98])).
cnf(c_0_105, plain, (v3_pre_topc(k3_xboole_0(X2,X3),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_pre_topc(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_106, negated_conjecture, (v3_pre_topc(esk4_0,esk1_0)|~r2_hidden(esk4_0,u1_struct_0(k9_yellow_6(esk1_0,X1)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_59]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_107, negated_conjecture, (v3_pre_topc(esk3_0,esk1_0)|v1_xboole_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_36])])).
cnf(c_0_108, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X4!=k3_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_109, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(esk3_0,u1_struct_0(k9_yellow_6(esk1_0,X1)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_65]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_110, negated_conjecture, (~v3_pre_topc(k1_setfam_1(k2_tarski(esk3_0,esk4_0)),esk1_0)|~r2_hidden(esk2_0,k1_setfam_1(k2_tarski(esk3_0,esk4_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_104]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_111, plain, (v3_pre_topc(k1_setfam_1(k2_tarski(X2,X3)),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v3_pre_topc(X3,X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(k3_yellow_1(u1_struct_0(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_105, c_0_34]), c_0_34]), c_0_23])).
cnf(c_0_112, negated_conjecture, (v3_pre_topc(esk4_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_68]), c_0_36])])).
cnf(c_0_113, negated_conjecture, (v3_pre_topc(esk3_0,esk1_0)), inference(sr,[status(thm)],[c_0_107, c_0_83])).
cnf(c_0_114, plain, (r2_hidden(X1,X4)|X4!=k1_setfam_1(k2_tarski(X2,X3))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_108, c_0_23])).
cnf(c_0_115, negated_conjecture, (r2_hidden(esk2_0,esk3_0)|v1_xboole_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_101]), c_0_36])])).
cnf(c_0_116, negated_conjecture, (~r2_hidden(esk2_0,k1_setfam_1(k2_tarski(esk3_0,esk4_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_112]), c_0_113]), c_0_59]), c_0_65]), c_0_37]), c_0_38])])).
cnf(c_0_117, plain, (r2_hidden(X1,k1_setfam_1(k2_tarski(X2,X3)))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_114])).
cnf(c_0_118, negated_conjecture, (r2_hidden(esk2_0,esk3_0)), inference(sr,[status(thm)],[c_0_115, c_0_83])).
cnf(c_0_119, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_73]), c_0_118])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 120
# Proof object clause steps            : 89
# Proof object formula steps           : 31
# Proof object conjectures             : 47
# Proof object clause conjectures      : 44
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 39
# Proof object simplifying inferences  : 83
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 36
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 33
# Processed clauses                    : 561
# ...of these trivial                  : 11
# ...subsumed                          : 242
# ...remaining for further processing  : 308
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 8
# Backward-rewritten                   : 52
# Generated clauses                    : 2971
# ...of the previous two non-trivial   : 2832
# Contextual simplify-reflections      : 7
# Paramodulations                      : 2936
# Factorizations                       : 24
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 206
#    Positive orientable unit clauses  : 33
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 10
#    Non-unit-clauses                  : 163
# Current number of unprocessed clauses: 2254
# ...number of literals in the above   : 10707
# Current number of archived formulas  : 0
# Current number of archived clauses   : 102
# Clause-clause subsumption calls (NU) : 9644
# Rec. Clause-clause subsumption calls : 5482
# Non-unit clause-clause subsumptions  : 183
# Unit Clause-clause subsumption calls : 613
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 35
# BW rewrite match successes           : 8
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 59444

# -------------------------------------------------
# User time                : 0.047 s
# System time              : 0.000 s
# Total time               : 0.047 s
# Maximum resident set size: 28688 pages
