# No SInE strategy applied
# Auto-Mode selected heuristic H_____047_B31_F1_PI_AE_R4_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t46_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0, file('data/MPTP/proofs/2f928e3a2178448dbdf792fc1e5d0876.E_input', t46_xboole_1)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/2f928e3a2178448dbdf792fc1e5d0876.E_input', t41_xboole_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/2f928e3a2178448dbdf792fc1e5d0876.E_input', t36_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/2f928e3a2178448dbdf792fc1e5d0876.E_input', t3_boole)).
fof(t50_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3)), file('data/MPTP/proofs/2f928e3a2178448dbdf792fc1e5d0876.E_input', t50_xboole_1)).
fof(t49_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),X3), file('data/MPTP/proofs/2f928e3a2178448dbdf792fc1e5d0876.E_input', t49_xboole_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/2f928e3a2178448dbdf792fc1e5d0876.E_input', t12_xboole_1)).
fof(t84_xboole_1, conjecture, ![X1, X2, X3]:~((~(r1_xboole_0(X1,X2))&(r1_xboole_0(X1,X3)&r1_xboole_0(X1,k4_xboole_0(X2,X3))))), file('data/MPTP/proofs/2f928e3a2178448dbdf792fc1e5d0876.E_input', t84_xboole_1)).
fof(t24_xboole_1, axiom, ![X1, X2, X3]:k2_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(k2_xboole_0(X1,X2),k2_xboole_0(X1,X3)), file('data/MPTP/proofs/2f928e3a2178448dbdf792fc1e5d0876.E_input', t24_xboole_1)).
fof(t21_xboole_1, axiom, ![X1, X2]:k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/2f928e3a2178448dbdf792fc1e5d0876.E_input', t21_xboole_1)).
fof(t78_xboole_1, axiom, ![X1, X2, X3]:(r1_xboole_0(X1,X2)=>k3_xboole_0(X1,k2_xboole_0(X2,X3))=k3_xboole_0(X1,X3)), file('data/MPTP/proofs/2f928e3a2178448dbdf792fc1e5d0876.E_input', t78_xboole_1)).
fof(t54_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(X1,k3_xboole_0(X2,X3))=k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,X3)), file('data/MPTP/proofs/2f928e3a2178448dbdf792fc1e5d0876.E_input', t54_xboole_1)).
fof(t43_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3)), file('data/MPTP/proofs/2f928e3a2178448dbdf792fc1e5d0876.E_input', t43_xboole_1)).
fof(commutativity_k5_xboole_0, axiom, ![X1, X2]:k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1), file('data/MPTP/proofs/2f928e3a2178448dbdf792fc1e5d0876.E_input', commutativity_k5_xboole_0)).
fof(d6_xboole_0, axiom, ![X1, X2]:k5_xboole_0(X1,X2)=k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X1)), file('data/MPTP/proofs/2f928e3a2178448dbdf792fc1e5d0876.E_input', d6_xboole_0)).
fof(t33_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(k4_xboole_0(X1,X3),k4_xboole_0(X2,X3))), file('data/MPTP/proofs/2f928e3a2178448dbdf792fc1e5d0876.E_input', t33_xboole_1)).
fof(t77_xboole_1, axiom, ![X1, X2, X3]:~((~(r1_xboole_0(X1,X2))&(r1_tarski(X1,X3)&r1_xboole_0(X1,k3_xboole_0(X2,X3))))), file('data/MPTP/proofs/2f928e3a2178448dbdf792fc1e5d0876.E_input', t77_xboole_1)).
fof(c_0_17, plain, ![X10, X11]:k4_xboole_0(X10,k2_xboole_0(X10,X11))=k1_xboole_0, inference(variable_rename,[status(thm)],[t46_xboole_1])).
fof(c_0_18, plain, ![X44, X45, X46]:k4_xboole_0(k4_xboole_0(X44,X45),X46)=k4_xboole_0(X44,k2_xboole_0(X45,X46)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
fof(c_0_19, plain, ![X37, X38]:r1_tarski(k4_xboole_0(X37,X38),X37), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_20, plain, ![X25]:k4_xboole_0(X25,k1_xboole_0)=X25, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_21, plain, ![X34, X35, X36]:k3_xboole_0(X34,k4_xboole_0(X35,X36))=k4_xboole_0(k3_xboole_0(X34,X35),k3_xboole_0(X34,X36)), inference(variable_rename,[status(thm)],[t50_xboole_1])).
fof(c_0_22, plain, ![X14, X15, X16]:k3_xboole_0(X14,k4_xboole_0(X15,X16))=k4_xboole_0(k3_xboole_0(X14,X15),X16), inference(variable_rename,[status(thm)],[t49_xboole_1])).
cnf(c_0_23, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_25, plain, ![X17, X18]:(~r1_tarski(X17,X18)|k2_xboole_0(X17,X18)=X18), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_26, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, plain, (k4_xboole_0(k4_xboole_0(X1,X1),X2)=k1_xboole_0), inference(rw,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_31, negated_conjecture, ~(![X1, X2, X3]:~((~(r1_xboole_0(X1,X2))&(r1_xboole_0(X1,X3)&r1_xboole_0(X1,k4_xboole_0(X2,X3)))))), inference(assume_negation,[status(cth)],[t84_xboole_1])).
fof(c_0_32, plain, ![X26, X27, X28]:k2_xboole_0(X26,k3_xboole_0(X27,X28))=k3_xboole_0(k2_xboole_0(X26,X27),k2_xboole_0(X26,X28)), inference(variable_rename,[status(thm)],[t24_xboole_1])).
cnf(c_0_33, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
fof(c_0_35, plain, ![X29, X30]:k3_xboole_0(X29,k2_xboole_0(X29,X30))=X29, inference(variable_rename,[status(thm)],[t21_xboole_1])).
cnf(c_0_36, plain, (k4_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3))=k4_xboole_0(k3_xboole_0(X1,X2),X3)), inference(rw,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_37, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_27, c_0_30])).
fof(c_0_38, plain, ![X19, X20, X21]:(~r1_xboole_0(X19,X20)|k3_xboole_0(X19,k2_xboole_0(X20,X21))=k3_xboole_0(X19,X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t78_xboole_1])])).
fof(c_0_39, negated_conjecture, (~r1_xboole_0(esk1_0,esk2_0)&(r1_xboole_0(esk1_0,esk3_0)&r1_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_31])])])])).
cnf(c_0_40, plain, (k2_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(k2_xboole_0(X1,X2),k2_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_41, plain, (k2_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_42, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_43, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X2)=k1_xboole_0), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_44, plain, (k3_xboole_0(X1,k2_xboole_0(X2,X3))=k3_xboole_0(X1,X3)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_45, negated_conjecture, (r1_xboole_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_46, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42])).
cnf(c_0_47, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_27, c_0_43])).
cnf(c_0_48, negated_conjecture, (k3_xboole_0(esk1_0,k2_xboole_0(esk3_0,X1))=k3_xboole_0(esk1_0,X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_49, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_50, plain, (k2_xboole_0(k4_xboole_0(X1,X2),X1)=X1), inference(spm,[status(thm)],[c_0_33, c_0_26])).
cnf(c_0_51, negated_conjecture, (r1_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_52, negated_conjecture, (k3_xboole_0(esk1_0,esk3_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_47])).
cnf(c_0_53, plain, (k3_xboole_0(k4_xboole_0(X1,X2),X1)=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_42, c_0_50])).
cnf(c_0_54, plain, (k4_xboole_0(k3_xboole_0(X1,k3_xboole_0(X2,X3)),k3_xboole_0(X2,X4))=k4_xboole_0(k3_xboole_0(X1,k3_xboole_0(X2,X3)),X4)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_36]), c_0_29])).
cnf(c_0_55, negated_conjecture, (k3_xboole_0(esk1_0,k2_xboole_0(k4_xboole_0(esk2_0,esk3_0),X1))=k3_xboole_0(esk1_0,X1)), inference(spm,[status(thm)],[c_0_44, c_0_51])).
cnf(c_0_56, negated_conjecture, (k4_xboole_0(k3_xboole_0(esk1_0,X1),esk3_0)=k3_xboole_0(esk1_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_52]), c_0_27])).
cnf(c_0_57, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_43, c_0_53])).
cnf(c_0_58, plain, (k4_xboole_0(k4_xboole_0(X1,X2),k3_xboole_0(X2,X3))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_24, c_0_46])).
cnf(c_0_59, plain, (k4_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_54, c_0_42])).
cnf(c_0_60, negated_conjecture, (k3_xboole_0(esk1_0,esk2_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_49]), c_0_29]), c_0_56]), c_0_47])).
fof(c_0_61, plain, ![X39, X40, X41]:k4_xboole_0(X39,k3_xboole_0(X40,X41))=k2_xboole_0(k4_xboole_0(X39,X40),k4_xboole_0(X39,X41)), inference(variable_rename,[status(thm)],[t54_xboole_1])).
cnf(c_0_62, plain, (k2_xboole_0(k1_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_50, c_0_37])).
cnf(c_0_63, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_64, negated_conjecture, (k4_xboole_0(k3_xboole_0(X1,esk1_0),esk2_0)=k3_xboole_0(X1,esk1_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_27])).
cnf(c_0_65, plain, (k4_xboole_0(X1,k3_xboole_0(X2,X3))=k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_66, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_42, c_0_62])).
fof(c_0_67, plain, ![X7, X8, X9]:(~r1_tarski(X7,k2_xboole_0(X8,X9))|r1_tarski(k4_xboole_0(X7,X8),X9)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
fof(c_0_68, plain, ![X42, X43]:k5_xboole_0(X42,X43)=k5_xboole_0(X43,X42), inference(variable_rename,[status(thm)],[commutativity_k5_xboole_0])).
fof(c_0_69, plain, ![X12, X13]:k5_xboole_0(X12,X13)=k2_xboole_0(k4_xboole_0(X12,X13),k4_xboole_0(X13,X12)), inference(variable_rename,[status(thm)],[d6_xboole_0])).
cnf(c_0_70, negated_conjecture, (k3_xboole_0(esk2_0,esk1_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_71, plain, (k2_xboole_0(X1,k4_xboole_0(X1,X2))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_27]), c_0_27])).
cnf(c_0_72, negated_conjecture, (k4_xboole_0(k3_xboole_0(esk1_0,X1),esk2_0)=k3_xboole_0(esk1_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_60]), c_0_27])).
cnf(c_0_73, plain, (k3_xboole_0(k2_xboole_0(X1,X2),X1)=k2_xboole_0(X1,k3_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_74, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_75, plain, (k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_76, plain, (k5_xboole_0(X1,X2)=k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_77, negated_conjecture, (k4_xboole_0(k3_xboole_0(esk2_0,X1),esk1_0)=k3_xboole_0(esk2_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_70]), c_0_27])).
cnf(c_0_78, negated_conjecture, (k4_xboole_0(k3_xboole_0(X1,esk2_0),esk1_0)=k3_xboole_0(X1,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_70]), c_0_27])).
cnf(c_0_79, plain, (k4_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X3))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_24, c_0_71])).
cnf(c_0_80, negated_conjecture, (k4_xboole_0(esk1_0,esk2_0)=esk1_0), inference(spm,[status(thm)],[c_0_72, c_0_42])).
fof(c_0_81, plain, ![X22, X23, X24]:(~r1_tarski(X22,X23)|r1_tarski(k4_xboole_0(X22,X24),k4_xboole_0(X23,X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t33_xboole_1])])).
cnf(c_0_82, plain, (k2_xboole_0(X1,k3_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_49]), c_0_47]), c_0_49]), c_0_73])).
cnf(c_0_83, plain, (r1_tarski(k4_xboole_0(k2_xboole_0(X1,X2),X1),X2)), inference(spm,[status(thm)],[c_0_74, c_0_34])).
cnf(c_0_84, plain, (k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X1))=k2_xboole_0(k4_xboole_0(X2,X1),k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_76]), c_0_76])).
cnf(c_0_85, negated_conjecture, (k4_xboole_0(esk2_0,esk1_0)=esk2_0), inference(spm,[status(thm)],[c_0_77, c_0_42])).
cnf(c_0_86, plain, (k4_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(k3_xboole_0(X1,X3),X4))=k4_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X3,X4))), inference(spm,[status(thm)],[c_0_36, c_0_29])).
cnf(c_0_87, negated_conjecture, (k4_xboole_0(k4_xboole_0(esk2_0,X1),esk1_0)=k4_xboole_0(esk2_0,X1)), inference(spm,[status(thm)],[c_0_78, c_0_53])).
cnf(c_0_88, negated_conjecture, (k4_xboole_0(esk1_0,k4_xboole_0(esk2_0,X1))=esk1_0), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_89, plain, (r1_tarski(k4_xboole_0(X1,X3),k4_xboole_0(X2,X3))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_90, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_24, c_0_41])).
cnf(c_0_91, plain, (k3_xboole_0(k2_xboole_0(X1,X2),X1)=X1), inference(rw,[status(thm)],[c_0_73, c_0_82])).
cnf(c_0_92, plain, (k2_xboole_0(k4_xboole_0(k2_xboole_0(X1,X2),X1),X2)=X2), inference(spm,[status(thm)],[c_0_33, c_0_83])).
cnf(c_0_93, negated_conjecture, (k2_xboole_0(esk1_0,esk2_0)=k2_xboole_0(esk2_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_80]), c_0_80])).
cnf(c_0_94, plain, (k4_xboole_0(X1,k4_xboole_0(k3_xboole_0(X1,X2),X3))=k4_xboole_0(X1,k4_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_86, c_0_42])).
fof(c_0_95, plain, ![X31, X32, X33]:(r1_xboole_0(X31,X32)|(~r1_tarski(X31,X33)|~r1_xboole_0(X31,k3_xboole_0(X32,X33)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t77_xboole_1])])])).
cnf(c_0_96, plain, (k3_xboole_0(X1,k2_xboole_0(k4_xboole_0(X1,X2),X3))=k2_xboole_0(k4_xboole_0(X1,X2),k3_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_40, c_0_50])).
cnf(c_0_97, negated_conjecture, (k2_xboole_0(k4_xboole_0(esk2_0,X1),esk1_0)=k2_xboole_0(esk1_0,k4_xboole_0(esk2_0,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_87]), c_0_88]), c_0_88])).
cnf(c_0_98, plain, (r1_tarski(k4_xboole_0(k4_xboole_0(X1,X2),X3),k4_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_89, c_0_26])).
cnf(c_0_99, plain, (k4_xboole_0(k4_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_90]), c_0_50])).
cnf(c_0_100, plain, (k4_xboole_0(k3_xboole_0(X1,k2_xboole_0(X2,X1)),X2)=k4_xboole_0(k2_xboole_0(X2,X1),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_29])).
cnf(c_0_101, negated_conjecture, (k3_xboole_0(esk1_0,k2_xboole_0(esk2_0,esk1_0))=esk1_0), inference(spm,[status(thm)],[c_0_42, c_0_93])).
cnf(c_0_102, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_63]), c_0_27])).
cnf(c_0_103, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,X3)|~r1_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_104, negated_conjecture, (k3_xboole_0(esk2_0,k2_xboole_0(esk1_0,k4_xboole_0(esk2_0,X1)))=k4_xboole_0(esk2_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_70]), c_0_49])).
cnf(c_0_105, plain, (r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X3,X2)))), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_106, negated_conjecture, (k4_xboole_0(k2_xboole_0(esk2_0,esk1_0),esk2_0)=esk1_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_100]), c_0_101])).
cnf(c_0_107, plain, (k2_xboole_0(k4_xboole_0(X1,X2),X2)=k2_xboole_0(X2,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_102]), c_0_90]), c_0_90])).
cnf(c_0_108, plain, (k4_xboole_0(k4_xboole_0(X1,k2_xboole_0(X2,X3)),X2)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_24]), c_0_29]), c_0_53])).
cnf(c_0_109, negated_conjecture, (r1_xboole_0(X1,esk2_0)|~r1_tarski(X1,k2_xboole_0(esk1_0,k4_xboole_0(esk2_0,X2)))|~r1_xboole_0(X1,k4_xboole_0(esk2_0,X2))), inference(spm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_110, negated_conjecture, (~r1_xboole_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_111, negated_conjecture, (r1_tarski(esk1_0,k2_xboole_0(k4_xboole_0(k2_xboole_0(esk2_0,esk1_0),X1),esk1_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_65]), c_0_106])).
cnf(c_0_112, plain, (k2_xboole_0(k4_xboole_0(X1,k2_xboole_0(X2,X3)),X2)=k2_xboole_0(X2,k4_xboole_0(X1,k2_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_113, negated_conjecture, (k4_xboole_0(k2_xboole_0(esk2_0,esk1_0),esk1_0)=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_100]), c_0_93]), c_0_93]), c_0_42])).
cnf(c_0_114, negated_conjecture, (~r1_tarski(esk1_0,k2_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_51]), c_0_110])).
cnf(c_0_115, negated_conjecture, (r1_tarski(esk1_0,k2_xboole_0(esk1_0,k4_xboole_0(esk2_0,X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_24]), c_0_113])).
cnf(c_0_116, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_114, c_0_115])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 117
# Proof object clause steps            : 82
# Proof object formula steps           : 35
# Proof object conjectures             : 31
# Proof object clause conjectures      : 28
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 58
# Proof object simplifying inferences  : 45
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 18
# Processed clauses                    : 5132
# ...of these trivial                  : 1666
# ...subsumed                          : 2234
# ...remaining for further processing  : 1232
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 7
# Backward-rewritten                   : 55
# Generated clauses                    : 192999
# ...of the previous two non-trivial   : 101383
# Contextual simplify-reflections      : 0
# Paramodulations                      : 192999
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1170
#    Positive orientable unit clauses  : 740
#    Positive unorientable unit clauses: 39
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 388
# Current number of unprocessed clauses: 96056
# ...number of literals in the above   : 115350
# Current number of archived formulas  : 0
# Current number of archived clauses   : 63
# Clause-clause subsumption calls (NU) : 14480
# Rec. Clause-clause subsumption calls : 14375
# Non-unit clause-clause subsumptions  : 2181
# Unit Clause-clause subsumption calls : 185
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 876
# BW rewrite match successes           : 117
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2280593

# -------------------------------------------------
# User time                : 1.073 s
# System time              : 0.037 s
# Total time               : 1.109 s
# Maximum resident set size: 85784 pages
