# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/2ff6473dfebd49c882ab65c7fc97f958.E_input', t5_subset)).
fof(dt_k2_orders_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k2_orders_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/2ff6473dfebd49c882ab65c7fc97f958.E_input', dt_k2_orders_2)).
fof(dt_k2_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/2ff6473dfebd49c882ab65c7fc97f958.E_input', dt_k2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/2ff6473dfebd49c882ab65c7fc97f958.E_input', dt_l1_orders_2)).
fof(d13_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k2_orders_2(X1,X2)=a_2_1_orders_2(X1,X2))), file('data/MPTP/proofs/2ff6473dfebd49c882ab65c7fc97f958.E_input', d13_orders_2)).
fof(d10_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_orders_2(X1,X2,X3)<=>(r1_orders_2(X1,X2,X3)&~(X2=X3)))))), file('data/MPTP/proofs/2ff6473dfebd49c882ab65c7fc97f958.E_input', d10_orders_2)).
fof(fraenkel_a_2_1_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&(v3_orders_2(X2)&(v4_orders_2(X2)&(v5_orders_2(X2)&l1_orders_2(X2)))))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r2_hidden(X1,a_2_1_orders_2(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r2_hidden(X5,X3)=>r2_orders_2(X2,X4,X5))))))), file('data/MPTP/proofs/2ff6473dfebd49c882ab65c7fc97f958.E_input', fraenkel_a_2_1_orders_2)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/2ff6473dfebd49c882ab65c7fc97f958.E_input', t4_subset)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/2ff6473dfebd49c882ab65c7fc97f958.E_input', t2_subset)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/2ff6473dfebd49c882ab65c7fc97f958.E_input', d3_struct_0)).
fof(t46_orders_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>k2_orders_2(X1,k2_struct_0(X1))=k1_xboole_0), file('data/MPTP/proofs/2ff6473dfebd49c882ab65c7fc97f958.E_input', t46_orders_2)).
fof(t3_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4]:((r2_hidden(X3,X4)&r2_hidden(X4,X2))=>r1_xboole_0(X3,X1)))))), file('data/MPTP/proofs/2ff6473dfebd49c882ab65c7fc97f958.E_input', t3_mcart_1)).
fof(c_0_12, plain, ![X15, X16, X17]:(~r2_hidden(X15,X16)|(~m1_subset_1(X16,k1_zfmisc_1(X17))|~v1_xboole_0(X17))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_13, plain, ![X22, X23]:(v2_struct_0(X22)|(~v3_orders_2(X22)|(~v4_orders_2(X22)|(~v5_orders_2(X22)|~l1_orders_2(X22))))|~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(X22)))|m1_subset_1(k2_orders_2(X22,X23),k1_zfmisc_1(u1_struct_0(X22)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_orders_2])])])).
cnf(c_0_14, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15, plain, (v2_struct_0(X1)|m1_subset_1(k2_orders_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_16, plain, ![X35]:(~l1_struct_0(X35)|m1_subset_1(k2_struct_0(X35),k1_zfmisc_1(u1_struct_0(X35)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_struct_0])])).
fof(c_0_17, plain, ![X7]:(~l1_orders_2(X7)|l1_struct_0(X7)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
fof(c_0_18, plain, ![X20, X21]:(v2_struct_0(X20)|(~v3_orders_2(X20)|(~v4_orders_2(X20)|(~v5_orders_2(X20)|~l1_orders_2(X20))))|(~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|k2_orders_2(X20,X21)=a_2_1_orders_2(X20,X21))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_orders_2])])])])).
fof(c_0_19, plain, ![X11, X12, X13]:(((r1_orders_2(X11,X12,X13)|~r2_orders_2(X11,X12,X13)|~m1_subset_1(X13,u1_struct_0(X11))|~m1_subset_1(X12,u1_struct_0(X11))|~l1_orders_2(X11))&(X12!=X13|~r2_orders_2(X11,X12,X13)|~m1_subset_1(X13,u1_struct_0(X11))|~m1_subset_1(X12,u1_struct_0(X11))|~l1_orders_2(X11)))&(~r1_orders_2(X11,X12,X13)|X12=X13|r2_orders_2(X11,X12,X13)|~m1_subset_1(X13,u1_struct_0(X11))|~m1_subset_1(X12,u1_struct_0(X11))|~l1_orders_2(X11))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_orders_2])])])])])).
fof(c_0_20, plain, ![X28, X29, X30, X32, X33]:(((m1_subset_1(esk3_3(X28,X29,X30),u1_struct_0(X29))|~r2_hidden(X28,a_2_1_orders_2(X29,X30))|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))))&((X28=esk3_3(X28,X29,X30)|~r2_hidden(X28,a_2_1_orders_2(X29,X30))|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))))&(~m1_subset_1(X32,u1_struct_0(X29))|(~r2_hidden(X32,X30)|r2_orders_2(X29,esk3_3(X28,X29,X30),X32))|~r2_hidden(X28,a_2_1_orders_2(X29,X30))|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))))))&((m1_subset_1(esk4_4(X28,X29,X30,X33),u1_struct_0(X29))|X28!=X33|~m1_subset_1(X33,u1_struct_0(X29))|r2_hidden(X28,a_2_1_orders_2(X29,X30))|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))))&((r2_hidden(esk4_4(X28,X29,X30,X33),X30)|X28!=X33|~m1_subset_1(X33,u1_struct_0(X29))|r2_hidden(X28,a_2_1_orders_2(X29,X30))|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))))&(~r2_orders_2(X29,X33,esk4_4(X28,X29,X30,X33))|X28!=X33|~m1_subset_1(X33,u1_struct_0(X29))|r2_hidden(X28,a_2_1_orders_2(X29,X30))|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_1_orders_2])])])])])])).
fof(c_0_21, plain, ![X8, X9, X10]:(~r2_hidden(X8,X9)|~m1_subset_1(X9,k1_zfmisc_1(X10))|m1_subset_1(X8,X10)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_22, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X3,k2_orders_2(X1,X2))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_23, plain, (m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (v2_struct_0(X1)|k2_orders_2(X1,X2)=a_2_1_orders_2(X1,X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, plain, (X1!=X2|~r2_orders_2(X3,X1,X2)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (r2_orders_2(X2,esk3_3(X4,X2,X3),X1)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r2_hidden(X4,a_2_1_orders_2(X2,X3))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_29, plain, ![X18, X19]:(~m1_subset_1(X18,X19)|(v1_xboole_0(X19)|r2_hidden(X18,X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_30, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~r2_hidden(X2,k2_orders_2(X1,k2_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24])).
cnf(c_0_31, plain, (k2_orders_2(X1,k2_struct_0(X1))=a_2_1_orders_2(X1,k2_struct_0(X1))|v2_struct_0(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_23]), c_0_24])).
fof(c_0_32, plain, ![X14]:(~l1_struct_0(X14)|k2_struct_0(X14)=u1_struct_0(X14)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
fof(c_0_33, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>k2_orders_2(X1,k2_struct_0(X1))=k1_xboole_0)), inference(assume_negation,[status(cth)],[t46_orders_2])).
cnf(c_0_34, plain, (~r2_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_26])).
cnf(c_0_35, plain, (r2_orders_2(X1,esk3_3(X2,X1,X3),X4)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,a_2_1_orders_2(X1,X3))|~r2_hidden(X4,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_36, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, plain, (m1_subset_1(esk3_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_orders_2(X2,X3))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_38, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~r2_hidden(X2,a_2_1_orders_2(X1,k2_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_39, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_40, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&k2_orders_2(esk1_0,k2_struct_0(esk1_0))!=k1_xboole_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_33])])])])).
cnf(c_0_41, plain, (v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(esk3_3(X3,X1,X2),X2)|~r2_hidden(X3,a_2_1_orders_2(X1,X2))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_28])).
cnf(c_0_42, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(esk3_3(X2,X1,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,a_2_1_orders_2(X1,X3))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_43, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~r2_hidden(X2,a_2_1_orders_2(X1,u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_24])).
cnf(c_0_44, negated_conjecture, (k2_orders_2(esk1_0,k2_struct_0(esk1_0))!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_45, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_46, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_47, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_48, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_49, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_50, plain, (v2_struct_0(X1)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,a_2_1_orders_2(X1,u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43])).
cnf(c_0_51, plain, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_23, c_0_39])).
fof(c_0_52, plain, ![X24, X26, X27]:((r2_hidden(esk2_1(X24),X24)|X24=k1_xboole_0)&(~r2_hidden(X26,X27)|~r2_hidden(X27,esk2_1(X24))|r1_xboole_0(X26,X24)|X24=k1_xboole_0)), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_mcart_1])])])])])])])).
cnf(c_0_53, negated_conjecture, (a_2_1_orders_2(esk1_0,k2_struct_0(esk1_0))!=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_31]), c_0_45]), c_0_46]), c_0_47]), c_0_48])]), c_0_49])).
cnf(c_0_54, plain, (v2_struct_0(X1)|~r2_hidden(X2,a_2_1_orders_2(X1,u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_24])).
cnf(c_0_55, plain, (r2_hidden(esk2_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_56, negated_conjecture, (a_2_1_orders_2(esk1_0,u1_struct_0(esk1_0))!=k1_xboole_0|~l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_53, c_0_39])).
cnf(c_0_57, plain, (a_2_1_orders_2(X1,u1_struct_0(X1))=k1_xboole_0|v2_struct_0(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_58, negated_conjecture, (~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_45]), c_0_46]), c_0_47]), c_0_48])]), c_0_49])).
cnf(c_0_59, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_24]), c_0_45])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 60
# Proof object clause steps            : 35
# Proof object formula steps           : 25
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 25
# Processed clauses                    : 145
# ...of these trivial                  : 0
# ...subsumed                          : 49
# ...remaining for further processing  : 96
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 8
# Backward-rewritten                   : 0
# Generated clauses                    : 207
# ...of the previous two non-trivial   : 202
# Contextual simplify-reflections      : 18
# Paramodulations                      : 203
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 84
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 75
# Current number of unprocessed clauses: 81
# ...number of literals in the above   : 803
# Current number of archived formulas  : 0
# Current number of archived clauses   : 8
# Clause-clause subsumption calls (NU) : 4922
# Rec. Clause-clause subsumption calls : 248
# Non-unit clause-clause subsumptions  : 72
# Unit Clause-clause subsumption calls : 119
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8839

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 28764 pages
