# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t53_zfmisc_1, conjecture, ![X1, X2, X3]:((r2_hidden(X1,X2)&r2_hidden(X3,X2))=>k3_xboole_0(k2_tarski(X1,X3),X2)=k2_tarski(X1,X3)), file('data/MPTP/proofs/3027ce1f04a74008b660fdc716139f12.E_input', t53_zfmisc_1)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/3027ce1f04a74008b660fdc716139f12.E_input', d4_xboole_0)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/3027ce1f04a74008b660fdc716139f12.E_input', t48_xboole_1)).
fof(d1_enumset1, axiom, ![X1, X2, X3, X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3)))))), file('data/MPTP/proofs/3027ce1f04a74008b660fdc716139f12.E_input', d1_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/3027ce1f04a74008b660fdc716139f12.E_input', t71_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/3027ce1f04a74008b660fdc716139f12.E_input', t70_enumset1)).
fof(c_0_6, negated_conjecture, ~(![X1, X2, X3]:((r2_hidden(X1,X2)&r2_hidden(X3,X2))=>k3_xboole_0(k2_tarski(X1,X3),X2)=k2_tarski(X1,X3))), inference(assume_negation,[status(cth)],[t53_zfmisc_1])).
fof(c_0_7, plain, ![X9, X10, X11, X12, X13, X14, X15, X16]:((((r2_hidden(X12,X9)|~r2_hidden(X12,X11)|X11!=k3_xboole_0(X9,X10))&(r2_hidden(X12,X10)|~r2_hidden(X12,X11)|X11!=k3_xboole_0(X9,X10)))&(~r2_hidden(X13,X9)|~r2_hidden(X13,X10)|r2_hidden(X13,X11)|X11!=k3_xboole_0(X9,X10)))&((~r2_hidden(esk4_3(X14,X15,X16),X16)|(~r2_hidden(esk4_3(X14,X15,X16),X14)|~r2_hidden(esk4_3(X14,X15,X16),X15))|X16=k3_xboole_0(X14,X15))&((r2_hidden(esk4_3(X14,X15,X16),X14)|r2_hidden(esk4_3(X14,X15,X16),X16)|X16=k3_xboole_0(X14,X15))&(r2_hidden(esk4_3(X14,X15,X16),X15)|r2_hidden(esk4_3(X14,X15,X16),X16)|X16=k3_xboole_0(X14,X15))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_8, plain, ![X18, X19]:k4_xboole_0(X18,k4_xboole_0(X18,X19))=k3_xboole_0(X18,X19), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_9, plain, ![X20, X21, X22, X23, X24, X25, X26, X27, X28, X29]:(((~r2_hidden(X24,X23)|(X24=X20|(X24=X21|X24=X22))|X23!=k1_enumset1(X20,X21,X22))&((X25!=X20|r2_hidden(X25,X23)|X23!=k1_enumset1(X20,X21,X22))&((X25!=X21|r2_hidden(X25,X23)|X23!=k1_enumset1(X20,X21,X22))&(X25!=X22|r2_hidden(X25,X23)|X23!=k1_enumset1(X20,X21,X22)))))&(((esk5_4(X26,X27,X28,X29)!=X26|~r2_hidden(esk5_4(X26,X27,X28,X29),X29)|X29=k1_enumset1(X26,X27,X28))&((esk5_4(X26,X27,X28,X29)!=X27|~r2_hidden(esk5_4(X26,X27,X28,X29),X29)|X29=k1_enumset1(X26,X27,X28))&(esk5_4(X26,X27,X28,X29)!=X28|~r2_hidden(esk5_4(X26,X27,X28,X29),X29)|X29=k1_enumset1(X26,X27,X28))))&(r2_hidden(esk5_4(X26,X27,X28,X29),X29)|(esk5_4(X26,X27,X28,X29)=X26|(esk5_4(X26,X27,X28,X29)=X27|esk5_4(X26,X27,X28,X29)=X28))|X29=k1_enumset1(X26,X27,X28)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])).
fof(c_0_10, plain, ![X31, X32, X33]:k2_enumset1(X31,X31,X32,X33)=k1_enumset1(X31,X32,X33), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_11, negated_conjecture, ((r2_hidden(esk1_0,esk2_0)&r2_hidden(esk3_0,esk2_0))&k3_xboole_0(k2_tarski(esk1_0,esk3_0),esk2_0)!=k2_tarski(esk1_0,esk3_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
fof(c_0_12, plain, ![X34, X35]:k1_enumset1(X34,X34,X35)=k2_tarski(X34,X35), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_13, plain, (r2_hidden(esk4_3(X1,X2,X3),X1)|r2_hidden(esk4_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15, plain, (X1=X3|X1=X4|X1=X5|~r2_hidden(X1,X2)|X2!=k1_enumset1(X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, negated_conjecture, (k3_xboole_0(k2_tarski(esk1_0,esk3_0),esk2_0)!=k2_tarski(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, plain, (r2_hidden(esk4_3(X1,X2,X3),X2)|r2_hidden(esk4_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_20, plain, (X3=k3_xboole_0(X1,X2)|~r2_hidden(esk4_3(X1,X2,X3),X3)|~r2_hidden(esk4_3(X1,X2,X3),X1)|~r2_hidden(esk4_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_21, plain, (X3=k4_xboole_0(X1,k4_xboole_0(X1,X2))|r2_hidden(esk4_3(X1,X2,X3),X3)|r2_hidden(esk4_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_22, plain, (X1=X5|X1=X4|X1=X3|X2!=k2_enumset1(X3,X3,X4,X5)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_23, negated_conjecture, (k4_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0),k4_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0),esk2_0))!=k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_18]), c_0_18]), c_0_16]), c_0_16]), c_0_14])).
cnf(c_0_24, plain, (X3=k4_xboole_0(X1,k4_xboole_0(X1,X2))|r2_hidden(esk4_3(X1,X2,X3),X3)|r2_hidden(esk4_3(X1,X2,X3),X2)), inference(rw,[status(thm)],[c_0_19, c_0_14])).
cnf(c_0_25, plain, (X3=k4_xboole_0(X1,k4_xboole_0(X1,X2))|~r2_hidden(esk4_3(X1,X2,X3),X3)|~r2_hidden(esk4_3(X1,X2,X3),X2)|~r2_hidden(esk4_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_20, c_0_14])).
cnf(c_0_26, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=X1|r2_hidden(esk4_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (X1=X2|X1=X3|X1=X4|~r2_hidden(X1,k2_enumset1(X4,X4,X3,X2))), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_28, negated_conjecture, (r2_hidden(esk4_3(k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0),esk2_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0)),k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0))|r2_hidden(esk4_3(k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0),esk2_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0)),esk2_0)), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24])])).
cnf(c_0_29, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X5,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_30, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=X1|~r2_hidden(esk4_3(X1,X2,X1),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_26])).
cnf(c_0_31, negated_conjecture, (esk4_3(k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0),esk2_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0))=esk3_0|esk4_3(k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0),esk2_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0))=esk1_0|r2_hidden(esk4_3(k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0),esk2_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0)),esk2_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_32, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_enumset1(X4,X4,X5,X2)), inference(rw,[status(thm)],[c_0_29, c_0_16])).
cnf(c_0_33, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X2,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_34, negated_conjecture, (esk4_3(k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0),esk2_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0))=esk1_0|esk4_3(k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0),esk2_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0))=esk3_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_23])).
cnf(c_0_35, plain, (r2_hidden(X1,k2_enumset1(X2,X2,X3,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_32])])).
cnf(c_0_36, negated_conjecture, (r2_hidden(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_37, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_enumset1(X2,X2,X4,X5)), inference(rw,[status(thm)],[c_0_33, c_0_16])).
cnf(c_0_38, negated_conjecture, (esk4_3(k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0),esk2_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0))=esk1_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_34]), c_0_35]), c_0_36])]), c_0_23])).
cnf(c_0_39, plain, (r2_hidden(X1,k2_enumset1(X1,X1,X2,X3))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_37])])).
cnf(c_0_40, negated_conjecture, (r2_hidden(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_41, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_38]), c_0_39]), c_0_40])]), c_0_23]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 42
# Proof object clause steps            : 29
# Proof object formula steps           : 13
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 27
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 17
# Processed clauses                    : 52
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 51
# Other redundant clauses eliminated   : 23
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 2
# Generated clauses                    : 198
# ...of the previous two non-trivial   : 161
# Contextual simplify-reflections      : 1
# Paramodulations                      : 173
# Factorizations                       : 5
# Equation resolutions                 : 23
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 24
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 16
# Current number of unprocessed clauses: 143
# ...number of literals in the above   : 724
# Current number of archived formulas  : 0
# Current number of archived clauses   : 23
# Clause-clause subsumption calls (NU) : 79
# Rec. Clause-clause subsumption calls : 53
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 7
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 13
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4928

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28792 pages
