# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t53_relset_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(~(v1_xboole_0(X3))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))=>![X5]:(m1_subset_1(X5,X1)=>(r2_hidden(X5,k8_relset_1(X1,X3,X4,X2))<=>?[X6]:(m1_subset_1(X6,X3)&(r2_hidden(k4_tarski(X5,X6),X4)&r2_hidden(X6,X2))))))))), file('data/MPTP/proofs/302bd3526b0d4aefa435c180c111c55a.E_input', t53_relset_1)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/302bd3526b0d4aefa435c180c111c55a.E_input', l3_subset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/302bd3526b0d4aefa435c180c111c55a.E_input', cc1_relset_1)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/302bd3526b0d4aefa435c180c111c55a.E_input', l54_zfmisc_1)).
fof(t166_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k10_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k2_relat_1(X3))&(r2_hidden(k4_tarski(X1,X4),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/302bd3526b0d4aefa435c180c111c55a.E_input', t166_relat_1)).
fof(redefinition_k8_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k8_relset_1(X1,X2,X3,X4)=k10_relat_1(X3,X4)), file('data/MPTP/proofs/302bd3526b0d4aefa435c180c111c55a.E_input', redefinition_k8_relset_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/302bd3526b0d4aefa435c180c111c55a.E_input', t1_subset)).
fof(t20_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)=>(r2_hidden(X1,k1_relat_1(X3))&r2_hidden(X2,k2_relat_1(X3))))), file('data/MPTP/proofs/302bd3526b0d4aefa435c180c111c55a.E_input', t20_relat_1)).
fof(c_0_8, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(~(v1_xboole_0(X3))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))=>![X5]:(m1_subset_1(X5,X1)=>(r2_hidden(X5,k8_relset_1(X1,X3,X4,X2))<=>?[X6]:(m1_subset_1(X6,X3)&(r2_hidden(k4_tarski(X5,X6),X4)&r2_hidden(X6,X2)))))))))), inference(assume_negation,[status(cth)],[t53_relset_1])).
fof(c_0_9, plain, ![X14, X15, X16]:(~m1_subset_1(X15,k1_zfmisc_1(X14))|(~r2_hidden(X16,X15)|r2_hidden(X16,X14))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
fof(c_0_10, negated_conjecture, ![X12]:(~v1_xboole_0(esk1_0)&(~v1_xboole_0(esk2_0)&(~v1_xboole_0(esk3_0)&(m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk3_0)))&(m1_subset_1(esk5_0,esk1_0)&((~r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0))|(~m1_subset_1(X12,esk3_0)|(~r2_hidden(k4_tarski(esk5_0,X12),esk4_0)|~r2_hidden(X12,esk2_0))))&((m1_subset_1(esk6_0,esk3_0)|r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0)))&((r2_hidden(k4_tarski(esk5_0,esk6_0),esk4_0)|r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0)))&(r2_hidden(esk6_0,esk2_0)|r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])])).
fof(c_0_11, plain, ![X17, X18, X19]:(~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(X17,X18)))|v1_relat_1(X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_12, plain, ![X20, X21, X22, X23]:(((r2_hidden(X20,X22)|~r2_hidden(k4_tarski(X20,X21),k2_zfmisc_1(X22,X23)))&(r2_hidden(X21,X23)|~r2_hidden(k4_tarski(X20,X21),k2_zfmisc_1(X22,X23))))&(~r2_hidden(X20,X22)|~r2_hidden(X21,X23)|r2_hidden(k4_tarski(X20,X21),k2_zfmisc_1(X22,X23)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
cnf(c_0_13, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (~r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0))|~m1_subset_1(X1,esk3_0)|~r2_hidden(k4_tarski(esk5_0,X1),esk4_0)|~r2_hidden(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, negated_conjecture, (r2_hidden(k4_tarski(esk5_0,esk6_0),esk4_0)|r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_17, plain, ![X24, X25, X26, X28]:(((r2_hidden(esk7_3(X24,X25,X26),k2_relat_1(X26))|~r2_hidden(X24,k10_relat_1(X26,X25))|~v1_relat_1(X26))&((r2_hidden(k4_tarski(X24,esk7_3(X24,X25,X26)),X26)|~r2_hidden(X24,k10_relat_1(X26,X25))|~v1_relat_1(X26))&(r2_hidden(esk7_3(X24,X25,X26),X25)|~r2_hidden(X24,k10_relat_1(X26,X25))|~v1_relat_1(X26))))&(~r2_hidden(X28,k2_relat_1(X26))|(~r2_hidden(k4_tarski(X24,X28),X26)|~r2_hidden(X28,X25))|r2_hidden(X24,k10_relat_1(X26,X25))|~v1_relat_1(X26))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t166_relat_1])])])])])).
cnf(c_0_18, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,X2))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(esk1_0,esk3_0))|~r2_hidden(X1,esk4_0)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
fof(c_0_21, plain, ![X29, X30, X31, X32]:(~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(X29,X30)))|k8_relset_1(X29,X30,X31,X32)=k10_relat_1(X31,X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_relset_1])])).
cnf(c_0_22, negated_conjecture, (r2_hidden(k4_tarski(esk5_0,esk6_0),esk4_0)|~r2_hidden(k4_tarski(esk5_0,X1),esk4_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk3_0)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_23, plain, (r2_hidden(k4_tarski(X1,esk7_3(X1,X2,X3)),X3)|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, negated_conjecture, (v1_relat_1(esk4_0)), inference(spm,[status(thm)],[c_0_18, c_0_14])).
fof(c_0_25, plain, ![X33, X34]:(~r2_hidden(X33,X34)|m1_subset_1(X33,X34)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_26, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(k4_tarski(X2,X1),esk4_0)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_27, plain, (k8_relset_1(X2,X3,X1,X4)=k10_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, negated_conjecture, (r2_hidden(esk6_0,esk2_0)|r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_29, plain, ![X35, X36, X37]:((r2_hidden(X35,k1_relat_1(X37))|~r2_hidden(k4_tarski(X35,X36),X37)|~v1_relat_1(X37))&(r2_hidden(X36,k2_relat_1(X37))|~r2_hidden(k4_tarski(X35,X36),X37)|~v1_relat_1(X37))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t20_relat_1])])])).
cnf(c_0_30, negated_conjecture, (r2_hidden(k4_tarski(esk5_0,esk6_0),esk4_0)|~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk2_0)|~r2_hidden(esk5_0,k10_relat_1(esk4_0,X1))|~m1_subset_1(esk7_3(esk5_0,X1,esk4_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24])])).
cnf(c_0_31, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, negated_conjecture, (r2_hidden(esk7_3(X1,X2,esk4_0),esk3_0)|~r2_hidden(X1,k10_relat_1(esk4_0,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_23]), c_0_24])])).
cnf(c_0_33, negated_conjecture, (k8_relset_1(esk1_0,esk3_0,esk4_0,X1)=k10_relat_1(esk4_0,X1)), inference(spm,[status(thm)],[c_0_27, c_0_14])).
cnf(c_0_34, negated_conjecture, (r2_hidden(esk6_0,esk2_0)|~r2_hidden(k4_tarski(esk5_0,X1),esk4_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk3_0)), inference(spm,[status(thm)],[c_0_15, c_0_28])).
cnf(c_0_35, negated_conjecture, (m1_subset_1(esk6_0,esk3_0)|r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_36, plain, (r2_hidden(X3,k10_relat_1(X2,X4))|~r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)|~r2_hidden(X1,X4)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_37, plain, (r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, negated_conjecture, (r2_hidden(k4_tarski(esk5_0,esk6_0),esk4_0)|~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk2_0)|~r2_hidden(esk5_0,k10_relat_1(esk4_0,X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])).
cnf(c_0_39, plain, (r2_hidden(esk7_3(X1,X2,X3),X2)|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_40, negated_conjecture, (r2_hidden(k4_tarski(esk5_0,esk6_0),esk4_0)|r2_hidden(esk5_0,k10_relat_1(esk4_0,esk2_0))), inference(rw,[status(thm)],[c_0_16, c_0_33])).
cnf(c_0_41, negated_conjecture, (r2_hidden(esk6_0,esk2_0)|~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk2_0)|~r2_hidden(esk5_0,k10_relat_1(esk4_0,X1))|~m1_subset_1(esk7_3(esk5_0,X1,esk4_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_23]), c_0_24])])).
cnf(c_0_42, negated_conjecture, (m1_subset_1(esk6_0,esk3_0)|~r2_hidden(k4_tarski(esk5_0,X1),esk4_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk3_0)), inference(spm,[status(thm)],[c_0_15, c_0_35])).
cnf(c_0_43, plain, (r2_hidden(X1,k10_relat_1(X2,X3))|~v1_relat_1(X2)|~r2_hidden(k4_tarski(X1,X4),X2)|~r2_hidden(X4,X3)), inference(csr,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_44, negated_conjecture, (r2_hidden(k4_tarski(esk5_0,esk6_0),esk4_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_24])]), c_0_40])).
cnf(c_0_45, negated_conjecture, (r2_hidden(esk6_0,esk2_0)|~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk2_0)|~r2_hidden(esk5_0,k10_relat_1(esk4_0,X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_31]), c_0_32])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk5_0,k10_relat_1(esk4_0,esk2_0))|r2_hidden(esk6_0,esk2_0)), inference(rw,[status(thm)],[c_0_28, c_0_33])).
cnf(c_0_47, negated_conjecture, (m1_subset_1(esk6_0,esk3_0)|~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk2_0)|~r2_hidden(esk5_0,k10_relat_1(esk4_0,X1))|~m1_subset_1(esk7_3(esk5_0,X1,esk4_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_23]), c_0_24])])).
cnf(c_0_48, negated_conjecture, (~r2_hidden(esk5_0,k10_relat_1(esk4_0,esk2_0))|~r2_hidden(k4_tarski(esk5_0,X1),esk4_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk3_0)), inference(rw,[status(thm)],[c_0_15, c_0_33])).
cnf(c_0_49, negated_conjecture, (r2_hidden(esk5_0,k10_relat_1(esk4_0,X1))|~r2_hidden(esk6_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_24])])).
cnf(c_0_50, negated_conjecture, (r2_hidden(esk6_0,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_39]), c_0_24])]), c_0_46])).
cnf(c_0_51, negated_conjecture, (m1_subset_1(esk6_0,esk3_0)|~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk2_0)|~r2_hidden(esk5_0,k10_relat_1(esk4_0,X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_31]), c_0_32])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk5_0,k10_relat_1(esk4_0,esk2_0))|m1_subset_1(esk6_0,esk3_0)), inference(rw,[status(thm)],[c_0_35, c_0_33])).
cnf(c_0_53, negated_conjecture, (~r2_hidden(k4_tarski(esk5_0,X1),esk4_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(esk6_0,esk3_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_39]), c_0_24])]), c_0_52])).
cnf(c_0_55, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_44]), c_0_50]), c_0_54])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 56
# Proof object clause steps            : 39
# Proof object formula steps           : 17
# Proof object conjectures             : 32
# Proof object clause conjectures      : 29
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 32
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 22
# Processed clauses                    : 132
# ...of these trivial                  : 0
# ...subsumed                          : 20
# ...remaining for further processing  : 112
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 22
# Generated clauses                    : 271
# ...of the previous two non-trivial   : 264
# Contextual simplify-reflections      : 9
# Paramodulations                      : 271
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 63
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 49
# Current number of unprocessed clauses: 175
# ...number of literals in the above   : 747
# Current number of archived formulas  : 0
# Current number of archived clauses   : 49
# Clause-clause subsumption calls (NU) : 1743
# Rec. Clause-clause subsumption calls : 1012
# Non-unit clause-clause subsumptions  : 34
# Unit Clause-clause subsumption calls : 22
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6547

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.000 s
# Total time               : 0.013 s
# Maximum resident set size: 28792 pages
