# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t22_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(m1_pre_topc(X2,X3)=>(~(r1_tsep_1(X2,X3))&~(r1_tsep_1(X3,X2))))))), file('data/MPTP/proofs/3068c1762e5b4f959677840b4b87998b.E_input', t22_tmap_1)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/3068c1762e5b4f959677840b4b87998b.E_input', dt_m1_pre_topc)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/3068c1762e5b4f959677840b4b87998b.E_input', dt_l1_pre_topc)).
fof(symmetry_r1_tsep_1, axiom, ![X1, X2]:((l1_struct_0(X1)&l1_struct_0(X2))=>(r1_tsep_1(X1,X2)=>r1_tsep_1(X2,X1))), file('data/MPTP/proofs/3068c1762e5b4f959677840b4b87998b.E_input', symmetry_r1_tsep_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/3068c1762e5b4f959677840b4b87998b.E_input', d10_xboole_0)).
fof(d9_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>(m1_pre_topc(X2,X1)<=>(r1_tarski(k2_struct_0(X2),k2_struct_0(X1))&![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>(r2_hidden(X3,u1_pre_topc(X2))<=>?[X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))&(r2_hidden(X4,u1_pre_topc(X1))&X3=k9_subset_1(u1_struct_0(X2),X4,k2_struct_0(X2)))))))))), file('data/MPTP/proofs/3068c1762e5b4f959677840b4b87998b.E_input', d9_pre_topc)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/3068c1762e5b4f959677840b4b87998b.E_input', d3_struct_0)).
fof(t68_xboole_1, axiom, ![X1, X2, X3]:(~(v1_xboole_0(X3))=>~((r1_tarski(X3,X1)&(r1_tarski(X3,X2)&r1_xboole_0(X1,X2))))), file('data/MPTP/proofs/3068c1762e5b4f959677840b4b87998b.E_input', t68_xboole_1)).
fof(d3_tsep_1, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(l1_struct_0(X2)=>(r1_tsep_1(X1,X2)<=>r1_xboole_0(u1_struct_0(X1),u1_struct_0(X2))))), file('data/MPTP/proofs/3068c1762e5b4f959677840b4b87998b.E_input', d3_tsep_1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/3068c1762e5b4f959677840b4b87998b.E_input', fc2_struct_0)).
fof(c_0_10, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(m1_pre_topc(X2,X3)=>(~(r1_tsep_1(X2,X3))&~(r1_tsep_1(X3,X2)))))))), inference(assume_negation,[status(cth)],[t22_tmap_1])).
fof(c_0_11, plain, ![X28, X29]:(~l1_pre_topc(X28)|(~m1_pre_topc(X29,X28)|l1_pre_topc(X29))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
fof(c_0_12, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk1_0))&(m1_pre_topc(esk2_0,esk3_0)&(r1_tsep_1(esk2_0,esk3_0)|r1_tsep_1(esk3_0,esk2_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_10])])])])).
fof(c_0_13, plain, ![X16]:(~l1_pre_topc(X16)|l1_struct_0(X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_14, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_18, plain, ![X11, X12]:(~l1_struct_0(X11)|~l1_struct_0(X12)|(~r1_tsep_1(X11,X12)|r1_tsep_1(X12,X11))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_tsep_1])])).
cnf(c_0_19, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, negated_conjecture, (l1_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])])).
cnf(c_0_21, negated_conjecture, (l1_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_17]), c_0_16])])).
fof(c_0_22, plain, ![X14, X15]:(((r1_tarski(X14,X15)|X14!=X15)&(r1_tarski(X15,X14)|X14!=X15))&(~r1_tarski(X14,X15)|~r1_tarski(X15,X14)|X14=X15)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_23, plain, ![X20, X21, X22, X24, X26]:(((r1_tarski(k2_struct_0(X21),k2_struct_0(X20))|~m1_pre_topc(X21,X20)|~l1_pre_topc(X21)|~l1_pre_topc(X20))&(((m1_subset_1(esk4_3(X20,X21,X22),k1_zfmisc_1(u1_struct_0(X20)))|~r2_hidden(X22,u1_pre_topc(X21))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))|~m1_pre_topc(X21,X20)|~l1_pre_topc(X21)|~l1_pre_topc(X20))&((r2_hidden(esk4_3(X20,X21,X22),u1_pre_topc(X20))|~r2_hidden(X22,u1_pre_topc(X21))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))|~m1_pre_topc(X21,X20)|~l1_pre_topc(X21)|~l1_pre_topc(X20))&(X22=k9_subset_1(u1_struct_0(X21),esk4_3(X20,X21,X22),k2_struct_0(X21))|~r2_hidden(X22,u1_pre_topc(X21))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))|~m1_pre_topc(X21,X20)|~l1_pre_topc(X21)|~l1_pre_topc(X20))))&(~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X20)))|(~r2_hidden(X24,u1_pre_topc(X20))|X22!=k9_subset_1(u1_struct_0(X21),X24,k2_struct_0(X21)))|r2_hidden(X22,u1_pre_topc(X21))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))|~m1_pre_topc(X21,X20)|~l1_pre_topc(X21)|~l1_pre_topc(X20))))&((m1_subset_1(esk5_2(X20,X21),k1_zfmisc_1(u1_struct_0(X21)))|~r1_tarski(k2_struct_0(X21),k2_struct_0(X20))|m1_pre_topc(X21,X20)|~l1_pre_topc(X21)|~l1_pre_topc(X20))&((~r2_hidden(esk5_2(X20,X21),u1_pre_topc(X21))|(~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X20)))|(~r2_hidden(X26,u1_pre_topc(X20))|esk5_2(X20,X21)!=k9_subset_1(u1_struct_0(X21),X26,k2_struct_0(X21))))|~r1_tarski(k2_struct_0(X21),k2_struct_0(X20))|m1_pre_topc(X21,X20)|~l1_pre_topc(X21)|~l1_pre_topc(X20))&((m1_subset_1(esk6_2(X20,X21),k1_zfmisc_1(u1_struct_0(X20)))|r2_hidden(esk5_2(X20,X21),u1_pre_topc(X21))|~r1_tarski(k2_struct_0(X21),k2_struct_0(X20))|m1_pre_topc(X21,X20)|~l1_pre_topc(X21)|~l1_pre_topc(X20))&((r2_hidden(esk6_2(X20,X21),u1_pre_topc(X20))|r2_hidden(esk5_2(X20,X21),u1_pre_topc(X21))|~r1_tarski(k2_struct_0(X21),k2_struct_0(X20))|m1_pre_topc(X21,X20)|~l1_pre_topc(X21)|~l1_pre_topc(X20))&(esk5_2(X20,X21)=k9_subset_1(u1_struct_0(X21),esk6_2(X20,X21),k2_struct_0(X21))|r2_hidden(esk5_2(X20,X21),u1_pre_topc(X21))|~r1_tarski(k2_struct_0(X21),k2_struct_0(X20))|m1_pre_topc(X21,X20)|~l1_pre_topc(X21)|~l1_pre_topc(X20))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_pre_topc])])])])])).
fof(c_0_24, plain, ![X8]:(~l1_struct_0(X8)|k2_struct_0(X8)=u1_struct_0(X8)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_25, plain, (r1_tsep_1(X2,X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~r1_tsep_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, negated_conjecture, (r1_tsep_1(esk2_0,esk3_0)|r1_tsep_1(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_27, negated_conjecture, (l1_struct_0(esk2_0)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_28, negated_conjecture, (l1_struct_0(esk3_0)), inference(spm,[status(thm)],[c_0_19, c_0_21])).
fof(c_0_29, plain, ![X17, X18, X19]:(v1_xboole_0(X19)|(~r1_tarski(X19,X17)|(~r1_tarski(X19,X18)|~r1_xboole_0(X17,X18)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t68_xboole_1])])])).
cnf(c_0_30, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, plain, (r1_tarski(k2_struct_0(X1),k2_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X1)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_33, plain, ![X9, X10]:((~r1_tsep_1(X9,X10)|r1_xboole_0(u1_struct_0(X9),u1_struct_0(X10))|~l1_struct_0(X10)|~l1_struct_0(X9))&(~r1_xboole_0(u1_struct_0(X9),u1_struct_0(X10))|r1_tsep_1(X9,X10)|~l1_struct_0(X10)|~l1_struct_0(X9))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tsep_1])])])])).
cnf(c_0_34, negated_conjecture, (r1_tsep_1(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_28])])).
cnf(c_0_35, plain, (v1_xboole_0(X1)|~r1_tarski(X1,X2)|~r1_tarski(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_30])).
cnf(c_0_37, plain, (r1_tarski(k2_struct_0(X1),k2_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[c_0_31, c_0_14])).
cnf(c_0_38, negated_conjecture, (m1_pre_topc(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_39, negated_conjecture, (k2_struct_0(esk2_0)=u1_struct_0(esk2_0)), inference(spm,[status(thm)],[c_0_32, c_0_27])).
cnf(c_0_40, negated_conjecture, (k2_struct_0(esk3_0)=u1_struct_0(esk3_0)), inference(spm,[status(thm)],[c_0_32, c_0_28])).
cnf(c_0_41, plain, (r1_xboole_0(u1_struct_0(X1),u1_struct_0(X2))|~r1_tsep_1(X1,X2)|~l1_struct_0(X2)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_34]), c_0_28]), c_0_27])])).
fof(c_0_43, plain, ![X13]:(v2_struct_0(X13)|~l1_struct_0(X13)|~v1_xboole_0(u1_struct_0(X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_44, plain, (v1_xboole_0(X1)|~r1_tarski(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_45, negated_conjecture, (r1_tarski(u1_struct_0(esk2_0),u1_struct_0(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39]), c_0_40]), c_0_21])])).
cnf(c_0_46, negated_conjecture, (r1_xboole_0(u1_struct_0(esk3_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_27]), c_0_28])])).
cnf(c_0_47, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_48, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])])).
cnf(c_0_49, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_50, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_27])]), c_0_49]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 51
# Proof object clause steps            : 30
# Proof object formula steps           : 21
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 24
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 85
# ...of these trivial                  : 2
# ...subsumed                          : 0
# ...remaining for further processing  : 83
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 69
# ...of the previous two non-trivial   : 65
# Contextual simplify-reflections      : 5
# Paramodulations                      : 66
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 51
#    Positive orientable unit clauses  : 22
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 26
# Current number of unprocessed clauses: 38
# ...number of literals in the above   : 170
# Current number of archived formulas  : 0
# Current number of archived clauses   : 30
# Clause-clause subsumption calls (NU) : 359
# Rec. Clause-clause subsumption calls : 55
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 114
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3878

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28688 pages
