# No SInE strategy applied
# Auto-Mode selected heuristic G_E___200_B02_F1_SE_CS_SP_PI_S2k
# and selection function SelectCQArNTEqFirst.
#
# Preprocessing time       : 0.010 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t62_funct_2, conjecture, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,k1_tarski(X1),X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_tarski(X1),X2)))))=>(~(X2=k1_xboole_0)=>k2_relset_1(k1_tarski(X1),X2,X3)=k1_tarski(k1_funct_1(X3,X1)))), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', t62_funct_2)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', t71_enumset1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', redefinition_k1_relset_1)).
fof(dt_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k1_relset_1(X1,X2,X3),k1_zfmisc_1(X1))), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', dt_k1_relset_1)).
fof(t142_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r1_tarski(X4,k1_enumset1(X1,X2,X3))<=>~((~(X4=k1_xboole_0)&(~(X4=k1_tarski(X1))&(~(X4=k1_tarski(X2))&(~(X4=k1_tarski(X3))&(~(X4=k2_tarski(X1,X2))&(~(X4=k2_tarski(X2,X3))&(~(X4=k2_tarski(X1,X3))&~(X4=k1_enumset1(X1,X2,X3))))))))))), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', t142_zfmisc_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', t3_subset)).
fof(t14_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(k1_relat_1(X2)=k1_tarski(X1)=>k2_relat_1(X2)=k1_tarski(k1_funct_1(X2,X1)))), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', t14_funct_1)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', redefinition_k2_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', cc1_relset_1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', t2_xboole_1)).
fof(d12_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2, X3]:(X3=k9_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(X5,k1_relat_1(X1))&(r2_hidden(X5,X2)&X4=k1_funct_1(X1,X5)))))), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', d12_funct_1)).
fof(s3_funct_1__e15_31__wellord2, axiom, ![X1]:?[X2]:((v1_relat_1(X2)&v1_funct_1(X2))&(k1_relat_1(X2)=X1&![X3]:(r2_hidden(X3,X1)=>k1_funct_1(X2,X3)=k1_tarski(X3)))), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', s3_funct_1__e15_31__wellord2)).
fof(t5_funct_2, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X3)=X1&![X4]:(r2_hidden(X4,X1)=>r2_hidden(k1_funct_1(X3,X4),X2)))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', t5_funct_2)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', t113_zfmisc_1)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', l3_subset_1)).
fof(cc3_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_funct_1(X2))), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', cc3_funct_1)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', t4_subset_1)).
fof(d4_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2, X3]:((r2_hidden(X2,k1_relat_1(X1))=>(X3=k1_funct_1(X1,X2)<=>r2_hidden(k4_tarski(X2,X3),X1)))&(~(r2_hidden(X2,k1_relat_1(X1)))=>(X3=k1_funct_1(X1,X2)<=>X3=k1_xboole_0)))), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', d4_funct_1)).
fof(t6_funct_2, axiom, ![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(r2_hidden(X3,X1)=>(X2=k1_xboole_0|r2_hidden(k1_funct_1(X4,X3),k2_relset_1(X1,X2,X4))))), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', t6_funct_2)).
fof(t60_relat_1, axiom, (k1_relat_1(k1_xboole_0)=k1_xboole_0&k2_relat_1(k1_xboole_0)=k1_xboole_0), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', t60_relat_1)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', l13_xboole_0)).
fof(fc2_xboole_0, axiom, ![X1]:~(v1_xboole_0(k1_tarski(X1))), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', fc2_xboole_0)).
fof(dt_o_0_0_xboole_0, axiom, v1_xboole_0(o_0_0_xboole_0), file('data/MPTP/proofs/30b4f2e1b2684885afbaa4081671f483.E_input', dt_o_0_0_xboole_0)).
fof(c_0_26, negated_conjecture, ~(![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,k1_tarski(X1),X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_tarski(X1),X2)))))=>(~(X2=k1_xboole_0)=>k2_relset_1(k1_tarski(X1),X2,X3)=k1_tarski(k1_funct_1(X3,X1))))), inference(assume_negation,[status(cth)],[t62_funct_2])).
fof(c_0_27, negated_conjecture, ((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,k1_tarski(esk1_0),esk2_0)&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k1_tarski(esk1_0),esk2_0)))))&(esk2_0!=k1_xboole_0&k2_relset_1(k1_tarski(esk1_0),esk2_0,esk3_0)!=k1_tarski(k1_funct_1(esk3_0,esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_26])])])])).
fof(c_0_28, plain, ![X72]:k2_tarski(X72,X72)=k1_tarski(X72), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_29, plain, ![X65, X66]:k1_enumset1(X65,X65,X66)=k2_tarski(X65,X66), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_30, plain, ![X19, X20, X21]:k2_enumset1(X19,X19,X20,X21)=k1_enumset1(X19,X20,X21), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_31, plain, ![X41, X42, X43]:(~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(X41,X42)))|k1_relset_1(X41,X42,X43)=k1_relat_1(X43)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_32, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k1_tarski(esk1_0),esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_35, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_36, plain, ![X29, X30, X31]:(~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(X29,X30)))|m1_subset_1(k1_relset_1(X29,X30,X31),k1_zfmisc_1(X29))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_relset_1])])).
cnf(c_0_37, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),esk2_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_33]), c_0_34]), c_0_35])).
fof(c_0_39, plain, ![X34, X35, X36, X37]:((~r1_tarski(X37,k1_enumset1(X34,X35,X36))|(X37=k1_xboole_0|(X37=k1_tarski(X34)|(X37=k1_tarski(X35)|(X37=k1_tarski(X36)|(X37=k2_tarski(X34,X35)|(X37=k2_tarski(X35,X36)|(X37=k2_tarski(X34,X36)|X37=k1_enumset1(X34,X35,X36)))))))))&((X37!=k1_xboole_0|r1_tarski(X37,k1_enumset1(X34,X35,X36)))&((X37!=k1_tarski(X34)|r1_tarski(X37,k1_enumset1(X34,X35,X36)))&((X37!=k1_tarski(X35)|r1_tarski(X37,k1_enumset1(X34,X35,X36)))&((X37!=k1_tarski(X36)|r1_tarski(X37,k1_enumset1(X34,X35,X36)))&((X37!=k2_tarski(X34,X35)|r1_tarski(X37,k1_enumset1(X34,X35,X36)))&((X37!=k2_tarski(X35,X36)|r1_tarski(X37,k1_enumset1(X34,X35,X36)))&((X37!=k2_tarski(X34,X36)|r1_tarski(X37,k1_enumset1(X34,X35,X36)))&(X37!=k1_enumset1(X34,X35,X36)|r1_tarski(X37,k1_enumset1(X34,X35,X36))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t142_zfmisc_1])])])])).
fof(c_0_40, plain, ![X32, X33]:((~m1_subset_1(X32,k1_zfmisc_1(X33))|r1_tarski(X32,X33))&(~r1_tarski(X32,X33)|m1_subset_1(X32,k1_zfmisc_1(X33)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_41, plain, (m1_subset_1(k1_relset_1(X2,X3,X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_42, negated_conjecture, (k1_relset_1(k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),esk2_0,esk3_0)=k1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
fof(c_0_43, plain, ![X44, X45]:(~v1_relat_1(X45)|~v1_funct_1(X45)|(k1_relat_1(X45)!=k1_tarski(X44)|k2_relat_1(X45)=k1_tarski(k1_funct_1(X45,X44)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t14_funct_1])])).
cnf(c_0_44, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|X1=k1_tarski(X3)|X1=k1_tarski(X4)|X1=k2_tarski(X2,X3)|X1=k2_tarski(X3,X4)|X1=k2_tarski(X2,X4)|X1=k1_enumset1(X2,X3,X4)|~r1_tarski(X1,k1_enumset1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_45, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(k1_relat_1(esk3_0),k1_zfmisc_1(k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_38]), c_0_42])).
fof(c_0_47, plain, ![X38, X39, X40]:(~m1_subset_1(X40,k1_zfmisc_1(k2_zfmisc_1(X38,X39)))|k2_relset_1(X38,X39,X40)=k2_relat_1(X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
cnf(c_0_48, plain, (k2_relat_1(X1)=k1_tarski(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|k1_relat_1(X1)!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49, plain, (X1=k1_xboole_0|X1=k2_enumset1(X4,X4,X4,X4)|X1=k2_enumset1(X3,X3,X3,X4)|X1=k2_enumset1(X3,X3,X3,X3)|X1=k2_enumset1(X2,X2,X3,X4)|X1=k2_enumset1(X2,X2,X2,X4)|X1=k2_enumset1(X2,X2,X2,X3)|X1=k2_enumset1(X2,X2,X2,X2)|~r1_tarski(X1,k2_enumset1(X2,X2,X3,X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_33]), c_0_33]), c_0_33]), c_0_34]), c_0_34]), c_0_34]), c_0_34]), c_0_34]), c_0_34]), c_0_35]), c_0_35]), c_0_35]), c_0_35]), c_0_35]), c_0_35]), c_0_35]), c_0_35])).
cnf(c_0_50, negated_conjecture, (r1_tarski(k1_relat_1(esk3_0),k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
fof(c_0_51, plain, ![X69, X70, X71]:(~m1_subset_1(X71,k1_zfmisc_1(k2_zfmisc_1(X69,X70)))|v1_relat_1(X71)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_52, negated_conjecture, (k2_relset_1(k1_tarski(esk1_0),esk2_0,esk3_0)!=k1_tarski(k1_funct_1(esk3_0,esk1_0))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_53, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_54, plain, ![X67, X68]:(~r2_hidden(X67,X68)|~r1_tarski(X68,X67)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_55, plain, ![X18]:r1_tarski(k1_xboole_0,X18), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_56, plain, ![X46, X47, X48, X49, X51, X52, X53, X54, X56]:((((r2_hidden(esk5_4(X46,X47,X48,X49),k1_relat_1(X46))|~r2_hidden(X49,X48)|X48!=k9_relat_1(X46,X47)|(~v1_relat_1(X46)|~v1_funct_1(X46)))&((r2_hidden(esk5_4(X46,X47,X48,X49),X47)|~r2_hidden(X49,X48)|X48!=k9_relat_1(X46,X47)|(~v1_relat_1(X46)|~v1_funct_1(X46)))&(X49=k1_funct_1(X46,esk5_4(X46,X47,X48,X49))|~r2_hidden(X49,X48)|X48!=k9_relat_1(X46,X47)|(~v1_relat_1(X46)|~v1_funct_1(X46)))))&(~r2_hidden(X52,k1_relat_1(X46))|(~r2_hidden(X52,X47)|X51!=k1_funct_1(X46,X52))|r2_hidden(X51,X48)|X48!=k9_relat_1(X46,X47)|(~v1_relat_1(X46)|~v1_funct_1(X46))))&((~r2_hidden(esk6_3(X46,X53,X54),X54)|(~r2_hidden(X56,k1_relat_1(X46))|(~r2_hidden(X56,X53)|esk6_3(X46,X53,X54)!=k1_funct_1(X46,X56)))|X54=k9_relat_1(X46,X53)|(~v1_relat_1(X46)|~v1_funct_1(X46)))&((r2_hidden(esk7_3(X46,X53,X54),k1_relat_1(X46))|r2_hidden(esk6_3(X46,X53,X54),X54)|X54=k9_relat_1(X46,X53)|(~v1_relat_1(X46)|~v1_funct_1(X46)))&((r2_hidden(esk7_3(X46,X53,X54),X53)|r2_hidden(esk6_3(X46,X53,X54),X54)|X54=k9_relat_1(X46,X53)|(~v1_relat_1(X46)|~v1_funct_1(X46)))&(esk6_3(X46,X53,X54)=k1_funct_1(X46,esk7_3(X46,X53,X54))|r2_hidden(esk6_3(X46,X53,X54),X54)|X54=k9_relat_1(X46,X53)|(~v1_relat_1(X46)|~v1_funct_1(X46))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_funct_1])])])])])])).
fof(c_0_57, plain, ![X58, X60]:((v1_relat_1(esk8_1(X58))&v1_funct_1(esk8_1(X58)))&(k1_relat_1(esk8_1(X58))=X58&(~r2_hidden(X60,X58)|k1_funct_1(esk8_1(X58),X60)=k1_tarski(X60)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[s3_funct_1__e15_31__wellord2])])])])).
fof(c_0_58, plain, ![X11, X12, X13]:(((v1_funct_1(X13)|(r2_hidden(esk4_3(X11,X12,X13),X11)|k1_relat_1(X13)!=X11)|(~v1_relat_1(X13)|~v1_funct_1(X13)))&((v1_funct_2(X13,X11,X12)|(r2_hidden(esk4_3(X11,X12,X13),X11)|k1_relat_1(X13)!=X11)|(~v1_relat_1(X13)|~v1_funct_1(X13)))&(m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12)))|(r2_hidden(esk4_3(X11,X12,X13),X11)|k1_relat_1(X13)!=X11)|(~v1_relat_1(X13)|~v1_funct_1(X13)))))&((v1_funct_1(X13)|(~r2_hidden(k1_funct_1(X13,esk4_3(X11,X12,X13)),X12)|k1_relat_1(X13)!=X11)|(~v1_relat_1(X13)|~v1_funct_1(X13)))&((v1_funct_2(X13,X11,X12)|(~r2_hidden(k1_funct_1(X13,esk4_3(X11,X12,X13)),X12)|k1_relat_1(X13)!=X11)|(~v1_relat_1(X13)|~v1_funct_1(X13)))&(m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12)))|(~r2_hidden(k1_funct_1(X13,esk4_3(X11,X12,X13)),X12)|k1_relat_1(X13)!=X11)|(~v1_relat_1(X13)|~v1_funct_1(X13)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_funct_2])])])])).
cnf(c_0_59, plain, (k2_relat_1(X1)=k2_enumset1(k1_funct_1(X1,X2),k1_funct_1(X1,X2),k1_funct_1(X1,X2),k1_funct_1(X1,X2))|k1_relat_1(X1)!=k2_enumset1(X2,X2,X2,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_33]), c_0_33]), c_0_34]), c_0_34]), c_0_35]), c_0_35])).
cnf(c_0_60, negated_conjecture, (k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)=k1_relat_1(esk3_0)|k1_relat_1(esk3_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_61, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_62, negated_conjecture, (k2_relset_1(k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),esk2_0,esk3_0)!=k2_enumset1(k1_funct_1(esk3_0,esk1_0),k1_funct_1(esk3_0,esk1_0),k1_funct_1(esk3_0,esk1_0),k1_funct_1(esk3_0,esk1_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_33]), c_0_33]), c_0_34]), c_0_34]), c_0_35]), c_0_35])).
cnf(c_0_63, negated_conjecture, (k2_relset_1(k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),esk2_0,esk3_0)=k2_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_53, c_0_38])).
fof(c_0_64, plain, ![X9, X10]:((k2_zfmisc_1(X9,X10)!=k1_xboole_0|(X9=k1_xboole_0|X10=k1_xboole_0))&((X9!=k1_xboole_0|k2_zfmisc_1(X9,X10)=k1_xboole_0)&(X10!=k1_xboole_0|k2_zfmisc_1(X9,X10)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_65, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_66, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_67, plain, (r2_hidden(esk7_3(X1,X2,X3),k1_relat_1(X1))|r2_hidden(esk6_3(X1,X2,X3),X3)|X3=k9_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_68, plain, (k1_relat_1(esk8_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_69, plain, (v1_relat_1(esk8_1(X1))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_70, plain, (v1_funct_1(esk8_1(X1))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_71, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|r2_hidden(esk4_3(X2,X3,X1),X2)|k1_relat_1(X1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_72, negated_conjecture, (k2_enumset1(k1_funct_1(X1,esk1_0),k1_funct_1(X1,esk1_0),k1_funct_1(X1,esk1_0),k1_funct_1(X1,esk1_0))=k2_relat_1(X1)|k1_relat_1(esk3_0)=k1_xboole_0|k1_relat_1(X1)!=k1_relat_1(esk3_0)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_73, negated_conjecture, (v1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_61, c_0_38])).
cnf(c_0_74, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_75, negated_conjecture, (k2_enumset1(k1_funct_1(esk3_0,esk1_0),k1_funct_1(esk3_0,esk1_0),k1_funct_1(esk3_0,esk1_0),k1_funct_1(esk3_0,esk1_0))!=k2_relat_1(esk3_0)), inference(rw,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_76, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_77, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_78, plain, (X1=k9_relat_1(esk8_1(X2),X3)|r2_hidden(esk6_3(esk8_1(X2),X3,X1),X1)|r2_hidden(esk7_3(esk8_1(X2),X3,X1),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_69]), c_0_70])])).
fof(c_0_79, plain, ![X24, X25, X26]:(~m1_subset_1(X25,k1_zfmisc_1(X24))|(~r2_hidden(X26,X25)|r2_hidden(X26,X24))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
cnf(c_0_80, plain, (r2_hidden(esk4_3(k1_relat_1(X1),X2,X1),k1_relat_1(X1))|m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),X2)))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(er,[status(thm)],[c_0_71])).
cnf(c_0_81, negated_conjecture, (k1_relat_1(esk3_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_72]), c_0_73]), c_0_74])]), c_0_75])).
cnf(c_0_82, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_76])).
cnf(c_0_83, plain, (X1=k9_relat_1(esk8_1(k1_xboole_0),X2)|r2_hidden(esk6_3(esk8_1(k1_xboole_0),X2,X1),X1)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
fof(c_0_84, plain, ![X27, X28]:(~v1_relat_1(X27)|~v1_funct_1(X27)|(~m1_subset_1(X28,k1_zfmisc_1(X27))|v1_funct_1(X28))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_funct_1])])])).
fof(c_0_85, plain, ![X22]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X22)), inference(variable_rename,[status(thm)],[t4_subset_1])).
cnf(c_0_86, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_87, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_82]), c_0_73]), c_0_74])]), c_0_77])).
cnf(c_0_88, plain, (k9_relat_1(esk8_1(k1_xboole_0),X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_77, c_0_83])).
fof(c_0_89, plain, ![X15, X16, X17]:(((X17!=k1_funct_1(X15,X16)|r2_hidden(k4_tarski(X16,X17),X15)|~r2_hidden(X16,k1_relat_1(X15))|(~v1_relat_1(X15)|~v1_funct_1(X15)))&(~r2_hidden(k4_tarski(X16,X17),X15)|X17=k1_funct_1(X15,X16)|~r2_hidden(X16,k1_relat_1(X15))|(~v1_relat_1(X15)|~v1_funct_1(X15))))&((X17!=k1_funct_1(X15,X16)|X17=k1_xboole_0|r2_hidden(X16,k1_relat_1(X15))|(~v1_relat_1(X15)|~v1_funct_1(X15)))&(X17!=k1_xboole_0|X17=k1_funct_1(X15,X16)|r2_hidden(X16,k1_relat_1(X15))|(~v1_relat_1(X15)|~v1_funct_1(X15))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_funct_1])])])])])).
cnf(c_0_90, plain, (v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_91, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_92, negated_conjecture, (v1_funct_2(esk3_0,k1_tarski(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_93, negated_conjecture, (~r2_hidden(X1,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_77])).
cnf(c_0_94, plain, (X1=k1_xboole_0|r2_hidden(esk6_3(esk8_1(k1_xboole_0),X2,X1),X1)), inference(rw,[status(thm)],[c_0_83, c_0_88])).
cnf(c_0_95, plain, (X1=k1_funct_1(X2,X3)|r2_hidden(X3,k1_relat_1(X2))|X1!=k1_xboole_0|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_96, plain, (v1_funct_1(k1_xboole_0)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_90, c_0_91])).
fof(c_0_97, plain, ![X61, X62, X63, X64]:(~v1_funct_1(X64)|(~v1_funct_2(X64,X61,X62)|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(X61,X62))))|(~r2_hidden(X63,X61)|(X62=k1_xboole_0|r2_hidden(k1_funct_1(X64,X63),k2_relset_1(X61,X62,X64))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_funct_2])])).
cnf(c_0_98, negated_conjecture, (v1_funct_2(esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_33]), c_0_34]), c_0_35])).
cnf(c_0_99, negated_conjecture, (esk3_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_100, plain, (k1_funct_1(X1,X2)=k1_xboole_0|r2_hidden(X2,k1_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(er,[status(thm)],[c_0_95])).
cnf(c_0_101, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
cnf(c_0_102, plain, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_61, c_0_91])).
cnf(c_0_103, plain, (v1_funct_1(k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_69]), c_0_70])])).
cnf(c_0_104, plain, (k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
fof(c_0_105, plain, ![X73]:(~v1_xboole_0(X73)|X73=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
fof(c_0_106, plain, ![X23]:~v1_xboole_0(k1_tarski(X23)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_xboole_0])])).
cnf(c_0_107, plain, (X3=k1_xboole_0|r2_hidden(k1_funct_1(X1,X4),k2_relset_1(X2,X3,X1))|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_108, negated_conjecture, (v1_funct_2(k1_xboole_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),esk2_0)), inference(rw,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_109, plain, (k1_funct_1(k1_xboole_0,X1)=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_102])]), c_0_103])]), c_0_77])).
cnf(c_0_110, plain, (k2_relset_1(X1,X2,k1_xboole_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_91]), c_0_104])).
cnf(c_0_111, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_112, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_113, plain, (v1_xboole_0(o_0_0_xboole_0)), inference(split_conjunct,[status(thm)],[dt_o_0_0_xboole_0])).
cnf(c_0_114, plain, (~v1_xboole_0(k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_115, negated_conjecture, (~r2_hidden(X1,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_109]), c_0_110]), c_0_91]), c_0_103])]), c_0_111]), c_0_77])).
cnf(c_0_116, plain, (o_0_0_xboole_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_112, c_0_113])).
cnf(c_0_117, plain, (~v1_xboole_0(k2_enumset1(X1,X1,X1,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_114, c_0_33]), c_0_34]), c_0_35])).
cnf(c_0_118, negated_conjecture, (k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_115, c_0_94])).
cnf(c_0_119, plain, (v1_xboole_0(k1_xboole_0)), inference(rw,[status(thm)],[c_0_113, c_0_116])).
cnf(c_0_120, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_119])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 121
# Proof object clause steps            : 70
# Proof object formula steps           : 51
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 33
# Proof object initial formulas used   : 26
# Proof object generating inferences   : 24
# Proof object simplifying inferences  : 76
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 26
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 60
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 55
# Processed clauses                    : 563
# ...of these trivial                  : 3
# ...subsumed                          : 241
# ...remaining for further processing  : 319
# Other redundant clauses eliminated   : 24
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 10
# Backward-rewritten                   : 129
# Generated clauses                    : 1922
# ...of the previous two non-trivial   : 1754
# Contextual simplify-reflections      : 2
# Paramodulations                      : 1897
# Factorizations                       : 0
# Equation resolutions                 : 26
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 159
#    Positive orientable unit clauses  : 36
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 21
#    Non-unit-clauses                  : 102
# Current number of unprocessed clauses: 1217
# ...number of literals in the above   : 5830
# Current number of archived formulas  : 0
# Current number of archived clauses   : 142
# Clause-clause subsumption calls (NU) : 3550
# Rec. Clause-clause subsumption calls : 1349
# Non-unit clause-clause subsumptions  : 174
# Unit Clause-clause subsumption calls : 439
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 97
# BW rewrite match successes           : 34
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 43588

# -------------------------------------------------
# User time                : 0.038 s
# System time              : 0.000 s
# Total time               : 0.038 s
# Maximum resident set size: 28616 pages
