# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', d1_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', d3_tarski)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', d10_xboole_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', fc1_xboole_0)).
fof(cc4_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', cc4_relset_1)).
fof(existence_m1_subset_1, axiom, ![X1]:?[X2]:m1_subset_1(X2,X1), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', existence_m1_subset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', t3_subset)).
fof(t6_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=>~((r2_hidden(X1,X4)&![X5, X6]:~((X1=k4_tarski(X5,X6)&(r2_hidden(X5,X2)&r2_hidden(X6,X3))))))), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', t6_relset_1)).
fof(d2_funct_2, axiom, ![X1, X2, X3]:(X3=k1_funct_2(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:((v1_relat_1(X5)&v1_funct_1(X5))&(X4=X5&(k1_relat_1(X5)=X1&r1_tarski(k2_relat_1(X5),X2)))))), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', d2_funct_2)).
fof(t121_funct_2, conjecture, ![X1, X2, X3]:(r2_hidden(X3,k1_funct_2(X1,X2))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', t121_funct_2)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', cc2_relset_1)).
fof(t3_funct_2, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_funct_1(X1)&(v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))))), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', t3_funct_2)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', t113_zfmisc_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', cc1_relset_1)).
fof(t172_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&(v5_relat_1(X2,X1)&v1_funct_1(X2)))=>![X3]:(r2_hidden(X3,k1_relat_1(X2))=>r2_hidden(k1_funct_1(X2,X3),X1))), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', t172_funct_1)).
fof(t4_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=>(r1_tarski(X1,X4)=>m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', t4_relset_1)).
fof(cc5_funct_2, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_funct_2(X3,X1,X2))=>(v1_funct_1(X3)&v1_partfun1(X3,X1))))), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', cc5_funct_2)).
fof(t11_relset_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>((r1_tarski(k1_relat_1(X3),X1)&r1_tarski(k2_relat_1(X3),X2))=>m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', t11_relset_1)).
fof(cc1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_partfun1(X3,X1))=>(v1_funct_1(X3)&v1_funct_2(X3,X1,X2)))), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', cc1_funct_2)).
fof(t5_funct_2, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X3)=X1&![X4]:(r2_hidden(X4,X1)=>r2_hidden(k1_funct_1(X3,X4),X2)))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))), file('data/MPTP/proofs/30c28fc1195a43cf97947830d50c3eb4.E_input', t5_funct_2)).
fof(c_0_20, plain, ![X7, X8, X9]:((~v1_xboole_0(X7)|~r2_hidden(X8,X7))&(r2_hidden(esk1_1(X9),X9)|v1_xboole_0(X9))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_21, plain, ![X11, X12, X13, X14, X15]:((~r1_tarski(X11,X12)|(~r2_hidden(X13,X11)|r2_hidden(X13,X12)))&((r2_hidden(esk2_2(X14,X15),X14)|r1_tarski(X14,X15))&(~r2_hidden(esk2_2(X14,X15),X15)|r1_tarski(X14,X15)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_22, plain, ![X17, X18]:(((r1_tarski(X17,X18)|X17!=X18)&(r1_tarski(X18,X17)|X17!=X18))&(~r1_tarski(X17,X18)|~r1_tarski(X18,X17)|X17=X18)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_23, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_25, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_27, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
fof(c_0_28, plain, ![X19, X20, X21]:(~v1_xboole_0(X19)|(~m1_subset_1(X21,k1_zfmisc_1(k2_zfmisc_1(X20,X19)))|v1_xboole_0(X21))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])).
fof(c_0_29, plain, ![X22]:m1_subset_1(esk3_1(X22),X22), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[existence_m1_subset_1])])).
cnf(c_0_30, plain, (X1=X2|~r1_tarski(X1,X2)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_31, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33, plain, (m1_subset_1(esk3_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34, plain, (k1_xboole_0=X1|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_35, plain, (v1_xboole_0(esk3_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
fof(c_0_36, plain, ![X24, X25]:((~m1_subset_1(X24,k1_zfmisc_1(X25))|r1_tarski(X24,X25))&(~r1_tarski(X24,X25)|m1_subset_1(X24,k1_zfmisc_1(X25)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_37, plain, ![X26, X27, X28, X29]:((X26=k4_tarski(esk4_4(X26,X27,X28,X29),esk5_4(X26,X27,X28,X29))|~r2_hidden(X26,X29)|~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(X27,X28))))&((r2_hidden(esk4_4(X26,X27,X28,X29),X27)|~r2_hidden(X26,X29)|~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(X27,X28))))&(r2_hidden(esk5_4(X26,X27,X28,X29),X28)|~r2_hidden(X26,X29)|~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(X27,X28)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_relset_1])])])])).
cnf(c_0_38, plain, (esk3_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k1_xboole_0|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_39, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_40, plain, (r2_hidden(esk5_4(X1,X2,X3,X4),X3)|~r2_hidden(X1,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_41, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_33, c_0_38])).
fof(c_0_42, plain, ![X32, X33, X34, X35, X37, X38, X39, X40, X41, X43]:(((((v1_relat_1(esk6_4(X32,X33,X34,X35))|~r2_hidden(X35,X34)|X34!=k1_funct_2(X32,X33))&(v1_funct_1(esk6_4(X32,X33,X34,X35))|~r2_hidden(X35,X34)|X34!=k1_funct_2(X32,X33)))&((X35=esk6_4(X32,X33,X34,X35)|~r2_hidden(X35,X34)|X34!=k1_funct_2(X32,X33))&((k1_relat_1(esk6_4(X32,X33,X34,X35))=X32|~r2_hidden(X35,X34)|X34!=k1_funct_2(X32,X33))&(r1_tarski(k2_relat_1(esk6_4(X32,X33,X34,X35)),X33)|~r2_hidden(X35,X34)|X34!=k1_funct_2(X32,X33)))))&(~v1_relat_1(X38)|~v1_funct_1(X38)|(X37!=X38|(k1_relat_1(X38)!=X32|~r1_tarski(k2_relat_1(X38),X33)))|r2_hidden(X37,X34)|X34!=k1_funct_2(X32,X33)))&((~r2_hidden(esk7_3(X39,X40,X41),X41)|(~v1_relat_1(X43)|~v1_funct_1(X43)|(esk7_3(X39,X40,X41)!=X43|(k1_relat_1(X43)!=X39|~r1_tarski(k2_relat_1(X43),X40))))|X41=k1_funct_2(X39,X40))&(((v1_relat_1(esk8_3(X39,X40,X41))|r2_hidden(esk7_3(X39,X40,X41),X41)|X41=k1_funct_2(X39,X40))&(v1_funct_1(esk8_3(X39,X40,X41))|r2_hidden(esk7_3(X39,X40,X41),X41)|X41=k1_funct_2(X39,X40)))&((esk7_3(X39,X40,X41)=esk8_3(X39,X40,X41)|r2_hidden(esk7_3(X39,X40,X41),X41)|X41=k1_funct_2(X39,X40))&((k1_relat_1(esk8_3(X39,X40,X41))=X39|r2_hidden(esk7_3(X39,X40,X41),X41)|X41=k1_funct_2(X39,X40))&(r1_tarski(k2_relat_1(esk8_3(X39,X40,X41)),X40)|r2_hidden(esk7_3(X39,X40,X41),X41)|X41=k1_funct_2(X39,X40))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_funct_2])])])])])])).
cnf(c_0_43, plain, (r1_tarski(esk3_1(k1_zfmisc_1(X1)),X1)), inference(spm,[status(thm)],[c_0_39, c_0_33])).
cnf(c_0_44, plain, (~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r2_hidden(X4,X1)|~v1_xboole_0(X3)), inference(spm,[status(thm)],[c_0_23, c_0_40])).
cnf(c_0_45, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(k2_zfmisc_1(X1,k1_xboole_0)))), inference(spm,[status(thm)],[c_0_41, c_0_27])).
cnf(c_0_46, plain, (k1_relat_1(esk6_4(X1,X2,X3,X4))=X1|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_47, plain, (X1=esk6_4(X2,X3,X4,X1)|~r2_hidden(X1,X4)|X4!=k1_funct_2(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_48, negated_conjecture, ~(![X1, X2, X3]:(r2_hidden(X3,k1_funct_2(X1,X2))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))), inference(assume_negation,[status(cth)],[t121_funct_2])).
cnf(c_0_49, plain, (v1_funct_1(esk6_4(X1,X2,X3,X4))|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_50, plain, (v1_relat_1(esk6_4(X1,X2,X3,X4))|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_51, plain, ![X48, X49, X50]:((v4_relat_1(X50,X48)|~m1_subset_1(X50,k1_zfmisc_1(k2_zfmisc_1(X48,X49))))&(v5_relat_1(X50,X49)|~m1_subset_1(X50,k1_zfmisc_1(k2_zfmisc_1(X48,X49))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_52, plain, (esk3_1(k1_zfmisc_1(X1))=X1|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_30, c_0_43])).
cnf(c_0_53, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_27])])).
cnf(c_0_54, plain, (r2_hidden(esk4_4(X1,X2,X3,X4),X2)|~r2_hidden(X1,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_55, plain, ![X51]:((v1_funct_1(X51)|(~v1_relat_1(X51)|~v1_funct_1(X51)))&((v1_funct_2(X51,k1_relat_1(X51),k2_relat_1(X51))|(~v1_relat_1(X51)|~v1_funct_1(X51)))&(m1_subset_1(X51,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X51),k2_relat_1(X51))))|(~v1_relat_1(X51)|~v1_funct_1(X51))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_funct_2])])])).
cnf(c_0_56, plain, (k1_relat_1(X1)=X2|X3!=k1_funct_2(X2,X4)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
fof(c_0_57, negated_conjecture, (r2_hidden(esk11_0,k1_funct_2(esk9_0,esk10_0))&(~v1_funct_1(esk11_0)|(~v1_funct_2(esk11_0,esk9_0,esk10_0)|~m1_subset_1(esk11_0,k1_zfmisc_1(k2_zfmisc_1(esk9_0,esk10_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_48])])])).
cnf(c_0_58, plain, (v1_funct_1(X1)|X2!=k1_funct_2(X3,X4)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_49, c_0_47])).
cnf(c_0_59, plain, (v1_relat_1(X1)|X2!=k1_funct_2(X3,X4)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_50, c_0_47])).
cnf(c_0_60, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_61, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_33, c_0_52])).
cnf(c_0_62, plain, (~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,X2)))|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_63, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_64, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_65, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_66, plain, (k1_relat_1(X1)=X2|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(er,[status(thm)],[c_0_56])).
cnf(c_0_67, negated_conjecture, (r2_hidden(esk11_0,k1_funct_2(esk9_0,esk10_0))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_68, plain, (v1_funct_1(X1)|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(er,[status(thm)],[c_0_58])).
cnf(c_0_69, plain, (v1_relat_1(X1)|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(er,[status(thm)],[c_0_59])).
cnf(c_0_70, plain, (~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r2_hidden(X4,X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_23, c_0_54])).
cnf(c_0_71, plain, (~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_xboole_0)))|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_53, c_0_40])).
cnf(c_0_72, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2)|~v1_xboole_0(k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_73, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_30, c_0_26])).
fof(c_0_74, plain, ![X52, X53]:((k2_zfmisc_1(X52,X53)!=k1_xboole_0|(X52=k1_xboole_0|X53=k1_xboole_0))&((X52!=k1_xboole_0|k2_zfmisc_1(X52,X53)=k1_xboole_0)&(X53!=k1_xboole_0|k2_zfmisc_1(X52,X53)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_75, plain, ![X54, X55, X56]:(~m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(X54,X55)))|v1_relat_1(X56)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_76, plain, (~r1_tarski(X1,k2_zfmisc_1(k1_xboole_0,X2))|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_77, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_64])).
cnf(c_0_78, plain, (r1_tarski(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_39, c_0_65])).
cnf(c_0_79, negated_conjecture, (k1_relat_1(esk11_0)=esk9_0), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_80, negated_conjecture, (v1_funct_1(esk11_0)), inference(spm,[status(thm)],[c_0_68, c_0_67])).
cnf(c_0_81, negated_conjecture, (v1_relat_1(esk11_0)), inference(spm,[status(thm)],[c_0_69, c_0_67])).
cnf(c_0_82, plain, (~r1_tarski(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X4,X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_70, c_0_63])).
cnf(c_0_83, plain, (~r1_tarski(X1,k2_zfmisc_1(X2,k1_xboole_0))|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_71, c_0_63])).
fof(c_0_84, plain, ![X57, X58, X59]:(~v1_relat_1(X58)|(~v5_relat_1(X58,X57)|~v1_funct_1(X58))|(~r2_hidden(X59,k1_relat_1(X58))|r2_hidden(k1_funct_1(X58,X59),X57))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t172_funct_1])])])).
cnf(c_0_85, plain, (v5_relat_1(X1,X2)|~v1_xboole_0(k2_zfmisc_1(X3,X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_86, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_87, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_88, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_74])).
fof(c_0_89, plain, ![X60, X61, X62, X63]:(~m1_subset_1(X63,k1_zfmisc_1(k2_zfmisc_1(X61,X62)))|(~r1_tarski(X60,X63)|m1_subset_1(X60,k1_zfmisc_1(k2_zfmisc_1(X61,X62))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_relset_1])])).
cnf(c_0_90, plain, (m1_subset_1(X1,X2)|~v1_xboole_0(esk3_1(X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_33, c_0_73])).
cnf(c_0_91, plain, (~r2_hidden(X1,k2_zfmisc_1(k1_xboole_0,X2))), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_92, plain, (r2_hidden(esk1_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_93, negated_conjecture, (r1_tarski(esk11_0,k2_zfmisc_1(esk9_0,k2_relat_1(esk11_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80]), c_0_81])])).
cnf(c_0_94, plain, (~r2_hidden(X1,k2_zfmisc_1(X2,X3))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_82, c_0_77])).
cnf(c_0_95, plain, (~r2_hidden(X1,k2_zfmisc_1(X2,k1_xboole_0))), inference(spm,[status(thm)],[c_0_83, c_0_77])).
cnf(c_0_96, plain, (r2_hidden(k1_funct_1(X1,X3),X2)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)|~v1_funct_1(X1)|~r2_hidden(X3,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_84])).
fof(c_0_97, plain, ![X64, X65, X66]:((v1_funct_1(X66)|(~v1_funct_1(X66)|~v1_funct_2(X66,X64,X65))|~m1_subset_1(X66,k1_zfmisc_1(k2_zfmisc_1(X64,X65)))|v1_xboole_0(X65))&(v1_partfun1(X66,X64)|(~v1_funct_1(X66)|~v1_funct_2(X66,X64,X65))|~m1_subset_1(X66,k1_zfmisc_1(k2_zfmisc_1(X64,X65)))|v1_xboole_0(X65))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_funct_2])])])])])).
cnf(c_0_98, plain, (v5_relat_1(X1,X2)|X3!=k1_xboole_0|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_27])])).
cnf(c_0_99, plain, (v1_relat_1(X1)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_87, c_0_63])).
cnf(c_0_100, plain, (v5_relat_1(X1,X2)|X3!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_60, c_0_86])).
cnf(c_0_101, plain, (v1_relat_1(X1)|X2!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_102, plain, (m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r1_tarski(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_103, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~v1_xboole_0(X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_90, c_0_52])).
cnf(c_0_104, plain, (v1_xboole_0(k2_zfmisc_1(k1_xboole_0,X1))), inference(spm,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_105, negated_conjecture, (k2_zfmisc_1(esk9_0,k2_relat_1(esk11_0))=esk11_0|~v1_xboole_0(k2_zfmisc_1(esk9_0,k2_relat_1(esk11_0)))), inference(spm,[status(thm)],[c_0_30, c_0_93])).
cnf(c_0_106, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_94, c_0_92])).
cnf(c_0_107, plain, (v1_xboole_0(k2_zfmisc_1(X1,k1_xboole_0))), inference(spm,[status(thm)],[c_0_95, c_0_92])).
cnf(c_0_108, plain, (~v5_relat_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X3,k1_relat_1(X1))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_23, c_0_96])).
cnf(c_0_109, plain, (v1_partfun1(X1,X2)|v1_xboole_0(X3)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_110, plain, (v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_111, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_25, c_0_31])).
cnf(c_0_112, plain, (~v5_relat_1(X1,k1_xboole_0)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(spm,[status(thm)],[c_0_53, c_0_96])).
cnf(c_0_113, plain, (v5_relat_1(X1,X2)|~v1_xboole_0(X1)), inference(er,[status(thm)],[c_0_98])).
cnf(c_0_114, plain, (v1_relat_1(X1)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_99, c_0_26])).
cnf(c_0_115, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(er,[status(thm)],[c_0_100])).
cnf(c_0_116, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(er,[status(thm)],[c_0_101])).
cnf(c_0_117, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r1_tarski(X1,X4)|~v1_xboole_0(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X4)), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_118, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_34, c_0_104])).
cnf(c_0_119, plain, (~r1_tarski(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X4,X1)|~v1_xboole_0(X3)), inference(spm,[status(thm)],[c_0_44, c_0_63])).
cnf(c_0_120, negated_conjecture, (k2_zfmisc_1(esk9_0,k2_relat_1(esk11_0))=esk11_0|~v1_xboole_0(esk9_0)), inference(spm,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_121, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_34, c_0_107])).
cnf(c_0_122, negated_conjecture, (~v5_relat_1(esk11_0,X1)|~r2_hidden(X2,esk9_0)|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_79]), c_0_80]), c_0_81])])).
cnf(c_0_123, plain, (v5_relat_1(X1,k2_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_60, c_0_65])).
cnf(c_0_124, plain, (v1_partfun1(X1,k1_relat_1(X1))|v1_xboole_0(k2_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_65]), c_0_110])).
cnf(c_0_125, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_126, plain, (esk3_1(k1_zfmisc_1(k1_xboole_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_111, c_0_43])).
cnf(c_0_127, plain, (~v1_funct_1(X1)|~r2_hidden(X2,k1_relat_1(X1))|~v1_xboole_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_114])).
fof(c_0_128, plain, ![X67, X68, X69]:(~v1_relat_1(X69)|(~r1_tarski(k1_relat_1(X69),X67)|~r1_tarski(k2_relat_1(X69),X68)|m1_subset_1(X69,k1_zfmisc_1(k2_zfmisc_1(X67,X68))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_relset_1])])).
cnf(c_0_129, plain, (r1_tarski(k2_relat_1(esk6_4(X1,X2,X3,X4)),X2)|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_130, plain, (~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~r2_hidden(X2,k1_relat_1(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_115]), c_0_116])).
cnf(c_0_131, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~r1_tarski(X1,X2)|~v1_xboole_0(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_27])])).
cnf(c_0_132, plain, (~r2_hidden(X1,k2_zfmisc_1(X2,X3))|~v1_xboole_0(X3)), inference(spm,[status(thm)],[c_0_119, c_0_77])).
cnf(c_0_133, negated_conjecture, (v1_xboole_0(esk11_0)|~v1_xboole_0(esk9_0)), inference(spm,[status(thm)],[c_0_106, c_0_120])).
cnf(c_0_134, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_121]), c_0_27])])).
cnf(c_0_135, negated_conjecture, (~r2_hidden(X1,esk9_0)|~v1_xboole_0(k2_relat_1(esk11_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_80]), c_0_81])])).
cnf(c_0_136, negated_conjecture, (v1_partfun1(esk11_0,esk9_0)|v1_xboole_0(k2_relat_1(esk11_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_79]), c_0_80]), c_0_81])])).
cnf(c_0_137, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,esk3_1(k1_zfmisc_1(X2)))), inference(spm,[status(thm)],[c_0_125, c_0_43])).
cnf(c_0_138, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_126]), c_0_27])])).
cnf(c_0_139, plain, (v1_xboole_0(k1_relat_1(X1))|~v1_funct_1(X1)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_127, c_0_92])).
cnf(c_0_140, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_128])).
cnf(c_0_141, plain, (r1_tarski(k2_relat_1(X1),X2)|X3!=k1_funct_2(X4,X2)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_129, c_0_47])).
cnf(c_0_142, plain, (v1_xboole_0(k1_relat_1(X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_130, c_0_92])).
cnf(c_0_143, negated_conjecture, (m1_subset_1(esk11_0,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(k2_zfmisc_1(esk9_0,k2_relat_1(esk11_0)))), inference(spm,[status(thm)],[c_0_131, c_0_93])).
cnf(c_0_144, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_132, c_0_92])).
cnf(c_0_145, negated_conjecture, (v1_xboole_0(esk11_0)|~m1_subset_1(esk9_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_133, c_0_134])).
cnf(c_0_146, negated_conjecture, (v1_partfun1(esk11_0,esk9_0)|~r2_hidden(X1,esk9_0)), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_147, plain, (r2_hidden(esk1_1(esk3_1(k1_zfmisc_1(X1))),X1)|v1_xboole_0(esk3_1(k1_zfmisc_1(X1)))), inference(spm,[status(thm)],[c_0_137, c_0_92])).
cnf(c_0_148, plain, (m1_subset_1(k1_relat_1(X1),k1_zfmisc_1(k1_xboole_0))|~v1_funct_1(X1)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_138, c_0_139])).
cnf(c_0_149, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),X2)))|~v1_relat_1(X1)|~r1_tarski(k2_relat_1(X1),X2)), inference(spm,[status(thm)],[c_0_140, c_0_77])).
cnf(c_0_150, plain, (r1_tarski(k2_relat_1(X1),X2)|~r2_hidden(X1,k1_funct_2(X3,X2))), inference(er,[status(thm)],[c_0_141])).
cnf(c_0_151, negated_conjecture, (v1_xboole_0(esk9_0)|~m1_subset_1(esk11_0,k1_zfmisc_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_79]), c_0_80])])).
cnf(c_0_152, negated_conjecture, (m1_subset_1(esk11_0,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(k2_relat_1(esk11_0))), inference(spm,[status(thm)],[c_0_143, c_0_144])).
cnf(c_0_153, negated_conjecture, (esk11_0=X1|~m1_subset_1(esk9_0,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_73, c_0_145])).
cnf(c_0_154, negated_conjecture, (v1_partfun1(esk11_0,esk9_0)|v1_xboole_0(esk3_1(k1_zfmisc_1(esk9_0)))), inference(spm,[status(thm)],[c_0_146, c_0_147])).
cnf(c_0_155, negated_conjecture, (m1_subset_1(esk9_0,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(esk11_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_79]), c_0_80])])).
cnf(c_0_156, negated_conjecture, (~v1_funct_1(esk11_0)|~v1_funct_2(esk11_0,esk9_0,esk10_0)|~m1_subset_1(esk11_0,k1_zfmisc_1(k2_zfmisc_1(esk9_0,esk10_0)))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_157, negated_conjecture, (m1_subset_1(esk11_0,k1_zfmisc_1(k2_zfmisc_1(esk9_0,X1)))|~r1_tarski(k2_relat_1(esk11_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_79]), c_0_81])])).
cnf(c_0_158, negated_conjecture, (r1_tarski(k2_relat_1(esk11_0),esk10_0)), inference(spm,[status(thm)],[c_0_150, c_0_67])).
cnf(c_0_159, plain, (v5_relat_1(esk3_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2)), inference(spm,[status(thm)],[c_0_60, c_0_33])).
cnf(c_0_160, plain, (v1_relat_1(esk3_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))), inference(spm,[status(thm)],[c_0_87, c_0_33])).
cnf(c_0_161, negated_conjecture, (esk9_0=X1|~m1_subset_1(esk11_0,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_73, c_0_151])).
cnf(c_0_162, negated_conjecture, (v1_partfun1(esk11_0,esk9_0)|m1_subset_1(esk11_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_152, c_0_136])).
cnf(c_0_163, negated_conjecture, (esk3_1(k1_zfmisc_1(esk9_0))=esk11_0|v1_partfun1(esk11_0,esk9_0)|~m1_subset_1(esk9_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_153, c_0_154])).
cnf(c_0_164, negated_conjecture, (m1_subset_1(esk9_0,k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(esk11_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_155, c_0_134])).
fof(c_0_165, plain, ![X70, X71, X72]:((v1_funct_1(X72)|(~v1_funct_1(X72)|~v1_partfun1(X72,X70))|~m1_subset_1(X72,k1_zfmisc_1(k2_zfmisc_1(X70,X71))))&(v1_funct_2(X72,X70,X71)|(~v1_funct_1(X72)|~v1_partfun1(X72,X70))|~m1_subset_1(X72,k1_zfmisc_1(k2_zfmisc_1(X70,X71))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_2])])])).
cnf(c_0_166, negated_conjecture, (~v1_funct_2(esk11_0,esk9_0,esk10_0)|~m1_subset_1(esk11_0,k1_zfmisc_1(k2_zfmisc_1(esk9_0,esk10_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_156, c_0_80])])).
cnf(c_0_167, negated_conjecture, (m1_subset_1(esk11_0,k1_zfmisc_1(k2_zfmisc_1(esk9_0,esk10_0)))), inference(spm,[status(thm)],[c_0_157, c_0_158])).
fof(c_0_168, plain, ![X73, X74, X75]:(((v1_funct_1(X75)|(r2_hidden(esk12_3(X73,X74,X75),X73)|k1_relat_1(X75)!=X73)|(~v1_relat_1(X75)|~v1_funct_1(X75)))&((v1_funct_2(X75,X73,X74)|(r2_hidden(esk12_3(X73,X74,X75),X73)|k1_relat_1(X75)!=X73)|(~v1_relat_1(X75)|~v1_funct_1(X75)))&(m1_subset_1(X75,k1_zfmisc_1(k2_zfmisc_1(X73,X74)))|(r2_hidden(esk12_3(X73,X74,X75),X73)|k1_relat_1(X75)!=X73)|(~v1_relat_1(X75)|~v1_funct_1(X75)))))&((v1_funct_1(X75)|(~r2_hidden(k1_funct_1(X75,esk12_3(X73,X74,X75)),X74)|k1_relat_1(X75)!=X73)|(~v1_relat_1(X75)|~v1_funct_1(X75)))&((v1_funct_2(X75,X73,X74)|(~r2_hidden(k1_funct_1(X75,esk12_3(X73,X74,X75)),X74)|k1_relat_1(X75)!=X73)|(~v1_relat_1(X75)|~v1_funct_1(X75)))&(m1_subset_1(X75,k1_zfmisc_1(k2_zfmisc_1(X73,X74)))|(~r2_hidden(k1_funct_1(X75,esk12_3(X73,X74,X75)),X74)|k1_relat_1(X75)!=X73)|(~v1_relat_1(X75)|~v1_funct_1(X75)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_funct_2])])])])).
cnf(c_0_169, plain, (~v1_funct_1(k1_xboole_0)|~r2_hidden(X1,k1_relat_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_159]), c_0_121]), c_0_126]), c_0_160]), c_0_121]), c_0_126])])).
cnf(c_0_170, negated_conjecture, (esk3_1(k1_zfmisc_1(esk9_0))=esk9_0|v1_partfun1(esk11_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_154]), c_0_162])).
cnf(c_0_171, negated_conjecture, (esk3_1(k1_zfmisc_1(esk9_0))=esk11_0|v1_partfun1(esk11_0,esk9_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_164]), c_0_162])).
cnf(c_0_172, plain, (v1_funct_2(X1,X2,X3)|~v1_funct_1(X1)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_165])).
cnf(c_0_173, negated_conjecture, (~v1_funct_2(esk11_0,esk9_0,esk10_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_166, c_0_167])])).
cnf(c_0_174, plain, (v1_funct_2(X1,X2,X3)|r2_hidden(esk12_3(X2,X3,X1),X2)|k1_relat_1(X1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_168])).
cnf(c_0_175, plain, (v1_xboole_0(k1_relat_1(k1_xboole_0))|~v1_funct_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_169, c_0_92])).
cnf(c_0_176, negated_conjecture, (esk3_1(k1_zfmisc_1(esk9_0))=k1_xboole_0|v1_partfun1(esk11_0,esk9_0)), inference(spm,[status(thm)],[c_0_34, c_0_154])).
cnf(c_0_177, negated_conjecture, (esk9_0=esk11_0|v1_partfun1(esk11_0,esk9_0)), inference(spm,[status(thm)],[c_0_170, c_0_171])).
cnf(c_0_178, negated_conjecture, (~v1_partfun1(esk11_0,esk9_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_167]), c_0_80])]), c_0_173])).
cnf(c_0_179, plain, (v1_funct_2(X1,k1_relat_1(X1),X2)|r2_hidden(esk12_3(k1_relat_1(X1),X2,X1),k1_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_174])).
cnf(c_0_180, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0|~v1_funct_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_34, c_0_175])).
cnf(c_0_181, plain, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_114, c_0_27])).
cnf(c_0_182, negated_conjecture, (esk9_0=k1_xboole_0|v1_partfun1(esk11_0,esk9_0)), inference(spm,[status(thm)],[c_0_176, c_0_170])).
cnf(c_0_183, negated_conjecture, (esk9_0=esk11_0), inference(sr,[status(thm)],[c_0_177, c_0_178])).
cnf(c_0_184, plain, (v1_partfun1(X1,k1_xboole_0)|v1_xboole_0(X2)|~v1_funct_2(X1,k1_xboole_0,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_109, c_0_118])).
cnf(c_0_185, plain, (v1_funct_2(k1_xboole_0,k1_xboole_0,X1)|~v1_funct_1(k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_180]), c_0_181])]), c_0_53])).
cnf(c_0_186, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_33, c_0_126])).
cnf(c_0_187, negated_conjecture, (esk11_0=k1_xboole_0), inference(rw,[status(thm)],[inference(sr,[status(thm)],[c_0_182, c_0_178]), c_0_183])).
cnf(c_0_188, plain, (v1_partfun1(k1_xboole_0,k1_xboole_0)|v1_xboole_0(X1)|~v1_funct_1(k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_185]), c_0_186])])).
cnf(c_0_189, negated_conjecture, (v1_funct_1(k1_xboole_0)), inference(rw,[status(thm)],[c_0_80, c_0_187])).
cnf(c_0_190, negated_conjecture, (~v1_partfun1(k1_xboole_0,k1_xboole_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_178, c_0_183]), c_0_187]), c_0_187])).
cnf(c_0_191, plain, (v1_xboole_0(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_188, c_0_189])]), c_0_190])).
cnf(c_0_192, negated_conjecture, (r2_hidden(k1_xboole_0,k1_funct_2(k1_xboole_0,esk10_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_183]), c_0_187]), c_0_187])).
cnf(c_0_193, plain, (~r2_hidden(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_191])])).
cnf(c_0_194, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_192, c_0_193]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 195
# Proof object clause steps            : 155
# Proof object formula steps           : 40
# Proof object conjectures             : 44
# Proof object clause conjectures      : 41
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 32
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 112
# Proof object simplifying inferences  : 68
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 20
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 52
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 47
# Processed clauses                    : 57715
# ...of these trivial                  : 389
# ...subsumed                          : 52690
# ...remaining for further processing  : 4636
# Other redundant clauses eliminated   : 947
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 390
# Backward-rewritten                   : 3530
# Generated clauses                    : 496752
# ...of the previous two non-trivial   : 434769
# Contextual simplify-reflections      : 475
# Paramodulations                      : 495454
# Factorizations                       : 0
# Equation resolutions                 : 1203
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 572
#    Positive orientable unit clauses  : 25
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 542
# Current number of unprocessed clauses: 375633
# ...number of literals in the above   : 1957039
# Current number of archived formulas  : 0
# Current number of archived clauses   : 4061
# Clause-clause subsumption calls (NU) : 5491139
# Rec. Clause-clause subsumption calls : 2769060
# Non-unit clause-clause subsumptions  : 44789
# Unit Clause-clause subsumption calls : 17657
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 278
# BW rewrite match successes           : 265
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6872917

# -------------------------------------------------
# User time                : 5.876 s
# System time              : 0.090 s
# Total time               : 5.966 s
# Maximum resident set size: 265820 pages
