# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t33_mcart_1, conjecture, ![X1, X2, X3, X4, X5, X6, X7, X8]:(k4_mcart_1(X1,X2,X3,X4)=k4_mcart_1(X5,X6,X7,X8)=>(X1=X5&(X2=X6&(X3=X7&X4=X8)))), file('data/MPTP/proofs/31344c270b2143c5a788225d2b0baf40.E_input', t33_mcart_1)).
fof(d4_mcart_1, axiom, ![X1, X2, X3, X4]:k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4), file('data/MPTP/proofs/31344c270b2143c5a788225d2b0baf40.E_input', d4_mcart_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/31344c270b2143c5a788225d2b0baf40.E_input', d3_mcart_1)).
fof(t33_zfmisc_1, axiom, ![X1, X2, X3, X4]:(k4_tarski(X1,X2)=k4_tarski(X3,X4)=>(X1=X3&X2=X4)), file('data/MPTP/proofs/31344c270b2143c5a788225d2b0baf40.E_input', t33_zfmisc_1)).
fof(d1_mcart_1, axiom, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>![X2]:(X2=k1_mcart_1(X1)<=>![X3, X4]:(X1=k4_tarski(X3,X4)=>X2=X3))), file('data/MPTP/proofs/31344c270b2143c5a788225d2b0baf40.E_input', d1_mcart_1)).
fof(d2_mcart_1, axiom, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>![X2]:(X2=k2_mcart_1(X1)<=>![X3, X4]:(X1=k4_tarski(X3,X4)=>X2=X4))), file('data/MPTP/proofs/31344c270b2143c5a788225d2b0baf40.E_input', d2_mcart_1)).
fof(c_0_6, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6, X7, X8]:(k4_mcart_1(X1,X2,X3,X4)=k4_mcart_1(X5,X6,X7,X8)=>(X1=X5&(X2=X6&(X3=X7&X4=X8))))), inference(assume_negation,[status(cth)],[t33_mcart_1])).
fof(c_0_7, plain, ![X17, X18, X19, X20]:k4_mcart_1(X17,X18,X19,X20)=k4_tarski(k3_mcart_1(X17,X18,X19),X20), inference(variable_rename,[status(thm)],[d4_mcart_1])).
fof(c_0_8, plain, ![X21, X22, X23]:k3_mcart_1(X21,X22,X23)=k4_tarski(k4_tarski(X21,X22),X23), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_9, negated_conjecture, (k4_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=k4_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0)&(esk1_0!=esk5_0|(esk2_0!=esk6_0|(esk3_0!=esk7_0|esk4_0!=esk8_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
cnf(c_0_10, plain, (k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_12, plain, ![X24, X25, X26, X27]:((X24=X26|k4_tarski(X24,X25)!=k4_tarski(X26,X27))&(X25=X27|k4_tarski(X24,X25)!=k4_tarski(X26,X27))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t33_zfmisc_1])])])).
cnf(c_0_13, negated_conjecture, (k4_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=k4_mcart_1(esk5_0,esk6_0,esk7_0,esk8_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_10, c_0_11])).
fof(c_0_15, plain, ![X28, X29, X30, X31, X32, X33, X34]:((X31!=k1_mcart_1(X28)|(X28!=k4_tarski(X32,X33)|X31=X32)|X28!=k4_tarski(X29,X30))&((X28=k4_tarski(esk9_2(X28,X34),esk10_2(X28,X34))|X34=k1_mcart_1(X28)|X28!=k4_tarski(X29,X30))&(X34!=esk9_2(X28,X34)|X34=k1_mcart_1(X28)|X28!=k4_tarski(X29,X30)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_mcart_1])])])])])])).
cnf(c_0_16, plain, (X1=X2|k4_tarski(X3,X1)!=k4_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, negated_conjecture, (k4_tarski(k4_tarski(k4_tarski(esk5_0,esk6_0),esk7_0),esk8_0)=k4_tarski(k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0),esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_14]), c_0_14])).
fof(c_0_18, plain, ![X37, X38, X39, X40, X41, X42, X43]:((X40!=k2_mcart_1(X37)|(X37!=k4_tarski(X41,X42)|X40=X42)|X37!=k4_tarski(X38,X39))&((X37=k4_tarski(esk11_2(X37,X43),esk12_2(X37,X43))|X43=k2_mcart_1(X37)|X37!=k4_tarski(X38,X39))&(X43!=esk12_2(X37,X43)|X43=k2_mcart_1(X37)|X37!=k4_tarski(X38,X39)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_mcart_1])])])])])])).
cnf(c_0_19, plain, (X1=X3|X1!=k1_mcart_1(X2)|X2!=k4_tarski(X3,X4)|X2!=k4_tarski(X5,X6)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, negated_conjecture, (X1=esk8_0|k4_tarski(X2,X1)!=k4_tarski(k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0),esk4_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_21, plain, (X1=X4|X1!=k2_mcart_1(X2)|X2!=k4_tarski(X3,X4)|X2!=k4_tarski(X5,X6)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, plain, (k1_mcart_1(k4_tarski(X1,X2))=X3|k4_tarski(X1,X2)!=k4_tarski(X3,X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_19])])).
cnf(c_0_23, negated_conjecture, (esk4_0=esk8_0), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_24, plain, (k2_mcart_1(k4_tarski(X1,X2))=X3|k4_tarski(X1,X2)!=k4_tarski(X4,X3)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_21])])).
cnf(c_0_25, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_26, negated_conjecture, (k4_tarski(k4_tarski(k4_tarski(esk5_0,esk6_0),esk7_0),esk8_0)=k4_tarski(k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0),esk8_0)), inference(rw,[status(thm)],[c_0_17, c_0_23])).
cnf(c_0_27, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(er,[status(thm)],[c_0_24])).
cnf(c_0_28, negated_conjecture, (k4_tarski(k4_tarski(esk5_0,esk6_0),esk7_0)=k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_25])).
cnf(c_0_29, negated_conjecture, (esk3_0=esk7_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_27])).
cnf(c_0_30, negated_conjecture, (k4_tarski(k4_tarski(esk5_0,esk6_0),esk7_0)=k4_tarski(k4_tarski(esk1_0,esk2_0),esk7_0)), inference(rw,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_31, negated_conjecture, (esk1_0!=esk5_0|esk2_0!=esk6_0|esk3_0!=esk7_0|esk4_0!=esk8_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_32, negated_conjecture, (k4_tarski(esk5_0,esk6_0)=k4_tarski(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_30]), c_0_25])).
cnf(c_0_33, negated_conjecture, (esk5_0!=esk1_0|esk2_0!=esk6_0|esk3_0!=esk7_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_23])])).
cnf(c_0_34, negated_conjecture, (esk2_0=esk6_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_32]), c_0_27])).
cnf(c_0_35, negated_conjecture, (esk5_0!=esk1_0|esk2_0!=esk6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_29])])).
cnf(c_0_36, negated_conjecture, (k4_tarski(esk5_0,esk6_0)=k4_tarski(esk1_0,esk6_0)), inference(rw,[status(thm)],[c_0_32, c_0_34])).
cnf(c_0_37, negated_conjecture, (esk5_0!=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_34])])).
cnf(c_0_38, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_36]), c_0_25]), c_0_37]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 39
# Proof object clause steps            : 26
# Proof object formula steps           : 13
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 7
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 12
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 10
# Processed clauses                    : 45
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 42
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 14
# Generated clauses                    : 94
# ...of the previous two non-trivial   : 94
# Contextual simplify-reflections      : 0
# Paramodulations                      : 83
# Factorizations                       : 0
# Equation resolutions                 : 13
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 12
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 5
# Current number of unprocessed clauses: 66
# ...number of literals in the above   : 154
# Current number of archived formulas  : 0
# Current number of archived clauses   : 26
# Clause-clause subsumption calls (NU) : 65
# Rec. Clause-clause subsumption calls : 59
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 3
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1611

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.003 s
# Total time               : 0.009 s
# Maximum resident set size: 28852 pages
