# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t38_zfmisc_1, conjecture, ![X1, X2, X3]:(r1_tarski(k2_tarski(X1,X2),X3)<=>(r2_hidden(X1,X3)&r2_hidden(X2,X3))), file('data/MPTP/proofs/316afa05c2b04564b17182e1f9b5be27.E_input', t38_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/316afa05c2b04564b17182e1f9b5be27.E_input', t70_enumset1)).
fof(d1_enumset1, axiom, ![X1, X2, X3, X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3)))))), file('data/MPTP/proofs/316afa05c2b04564b17182e1f9b5be27.E_input', d1_enumset1)).
fof(t25_zfmisc_1, axiom, ![X1, X2, X3]:~((r1_tarski(k1_tarski(X1),k2_tarski(X2,X3))&(~(X1=X2)&~(X1=X3)))), file('data/MPTP/proofs/316afa05c2b04564b17182e1f9b5be27.E_input', t25_zfmisc_1)).
fof(l1_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/316afa05c2b04564b17182e1f9b5be27.E_input', l1_zfmisc_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/316afa05c2b04564b17182e1f9b5be27.E_input', d3_tarski)).
fof(t11_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k2_xboole_0(X1,X2),X3)=>r1_tarski(X1,X3)), file('data/MPTP/proofs/316afa05c2b04564b17182e1f9b5be27.E_input', t11_xboole_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/316afa05c2b04564b17182e1f9b5be27.E_input', t12_xboole_1)).
fof(t12_zfmisc_1, axiom, ![X1, X2]:r1_tarski(k1_tarski(X1),k2_tarski(X1,X2)), file('data/MPTP/proofs/316afa05c2b04564b17182e1f9b5be27.E_input', t12_zfmisc_1)).
fof(t10_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(X1,k2_xboole_0(X3,X2))), file('data/MPTP/proofs/316afa05c2b04564b17182e1f9b5be27.E_input', t10_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/316afa05c2b04564b17182e1f9b5be27.E_input', d10_xboole_0)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/316afa05c2b04564b17182e1f9b5be27.E_input', t7_xboole_1)).
fof(t8_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X3,X2))=>r1_tarski(k2_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/316afa05c2b04564b17182e1f9b5be27.E_input', t8_xboole_1)).
fof(c_0_13, negated_conjecture, ~(![X1, X2, X3]:(r1_tarski(k2_tarski(X1,X2),X3)<=>(r2_hidden(X1,X3)&r2_hidden(X2,X3)))), inference(assume_negation,[status(cth)],[t38_zfmisc_1])).
fof(c_0_14, negated_conjecture, ((~r1_tarski(k2_tarski(esk1_0,esk2_0),esk3_0)|(~r2_hidden(esk1_0,esk3_0)|~r2_hidden(esk2_0,esk3_0)))&((r2_hidden(esk1_0,esk3_0)|r1_tarski(k2_tarski(esk1_0,esk2_0),esk3_0))&(r2_hidden(esk2_0,esk3_0)|r1_tarski(k2_tarski(esk1_0,esk2_0),esk3_0)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])).
fof(c_0_15, plain, ![X9, X10]:k1_enumset1(X9,X9,X10)=k2_tarski(X9,X10), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_16, plain, ![X11, X12, X13, X14, X15, X16, X17, X18, X19, X20]:(((~r2_hidden(X15,X14)|(X15=X11|(X15=X12|X15=X13))|X14!=k1_enumset1(X11,X12,X13))&((X16!=X11|r2_hidden(X16,X14)|X14!=k1_enumset1(X11,X12,X13))&((X16!=X12|r2_hidden(X16,X14)|X14!=k1_enumset1(X11,X12,X13))&(X16!=X13|r2_hidden(X16,X14)|X14!=k1_enumset1(X11,X12,X13)))))&(((esk4_4(X17,X18,X19,X20)!=X17|~r2_hidden(esk4_4(X17,X18,X19,X20),X20)|X20=k1_enumset1(X17,X18,X19))&((esk4_4(X17,X18,X19,X20)!=X18|~r2_hidden(esk4_4(X17,X18,X19,X20),X20)|X20=k1_enumset1(X17,X18,X19))&(esk4_4(X17,X18,X19,X20)!=X19|~r2_hidden(esk4_4(X17,X18,X19,X20),X20)|X20=k1_enumset1(X17,X18,X19))))&(r2_hidden(esk4_4(X17,X18,X19,X20),X20)|(esk4_4(X17,X18,X19,X20)=X17|(esk4_4(X17,X18,X19,X20)=X18|esk4_4(X17,X18,X19,X20)=X19))|X20=k1_enumset1(X17,X18,X19)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])).
fof(c_0_17, plain, ![X22, X23, X24]:(~r1_tarski(k1_tarski(X22),k2_tarski(X23,X24))|(X22=X23|X22=X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t25_zfmisc_1])])])).
fof(c_0_18, plain, ![X25, X26]:((~r1_tarski(k1_tarski(X25),X26)|r2_hidden(X25,X26))&(~r2_hidden(X25,X26)|r1_tarski(k1_tarski(X25),X26))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])).
fof(c_0_19, plain, ![X27, X28, X29, X30, X31]:((~r1_tarski(X27,X28)|(~r2_hidden(X29,X27)|r2_hidden(X29,X28)))&((r2_hidden(esk5_2(X30,X31),X30)|r1_tarski(X30,X31))&(~r2_hidden(esk5_2(X30,X31),X31)|r1_tarski(X30,X31)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_20, negated_conjecture, (r2_hidden(esk2_0,esk3_0)|r1_tarski(k2_tarski(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X5,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_23, plain, ![X33, X34, X35]:(~r1_tarski(k2_xboole_0(X33,X34),X35)|r1_tarski(X33,X35)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_xboole_1])])).
fof(c_0_24, plain, ![X36, X37]:(~r1_tarski(X36,X37)|k2_xboole_0(X36,X37)=X37), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_25, plain, (X1=X2|X1=X3|~r1_tarski(k1_tarski(X1),k2_tarski(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, negated_conjecture, (r2_hidden(esk2_0,esk3_0)|r1_tarski(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0)), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_30, plain, (r2_hidden(X1,X2)|X2!=k1_enumset1(X3,X4,X1)), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_31, plain, (r1_tarski(X1,X3)|~r1_tarski(k2_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, negated_conjecture, (r2_hidden(esk1_0,esk3_0)|r1_tarski(k2_tarski(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_34, plain, ![X38, X39]:r1_tarski(k1_tarski(X38),k2_tarski(X38,X39)), inference(variable_rename,[status(thm)],[t12_zfmisc_1])).
cnf(c_0_35, plain, (X1=X3|X1=X2|~r1_tarski(k1_tarski(X1),k1_enumset1(X2,X2,X3))), inference(rw,[status(thm)],[c_0_25, c_0_21])).
cnf(c_0_36, plain, (r1_tarski(k1_tarski(esk5_2(X1,X2)),X1)|r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
fof(c_0_37, plain, ![X40, X41, X42]:(~r1_tarski(X40,X41)|r1_tarski(X40,k2_xboole_0(X42,X41))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_xboole_1])])).
cnf(c_0_38, negated_conjecture, (r2_hidden(esk2_0,esk3_0)|r2_hidden(X1,esk3_0)|~r2_hidden(X1,k1_enumset1(esk1_0,esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_39, plain, (r2_hidden(X1,k1_enumset1(X2,X3,X1))), inference(er,[status(thm)],[c_0_30])).
fof(c_0_40, plain, ![X43, X44]:(((r1_tarski(X43,X44)|X43!=X44)&(r1_tarski(X44,X43)|X43!=X44))&(~r1_tarski(X43,X44)|~r1_tarski(X44,X43)|X43=X44)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_41, plain, ![X45, X46]:r1_tarski(X45,k2_xboole_0(X45,X46)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_42, plain, (r1_tarski(X1,X2)|~r1_tarski(X3,X2)|~r1_tarski(X1,X3)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_43, negated_conjecture, (r2_hidden(esk1_0,esk3_0)|r1_tarski(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0)), inference(rw,[status(thm)],[c_0_33, c_0_21])).
cnf(c_0_44, plain, (r1_tarski(k1_tarski(X1),k2_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45, plain, (r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_46, plain, (esk5_2(k1_enumset1(X1,X1,X2),X3)=X1|esk5_2(k1_enumset1(X1,X1,X2),X3)=X2|r1_tarski(k1_enumset1(X1,X1,X2),X3)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_47, plain, (r2_hidden(X1,X2)|~r1_tarski(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_48, plain, (r1_tarski(X1,k2_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_49, negated_conjecture, (r2_hidden(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_50, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_51, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_52, plain, ![X47, X48, X49]:(~r1_tarski(X47,X48)|~r1_tarski(X49,X48)|r1_tarski(k2_xboole_0(X47,X49),X48)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_xboole_1])])).
cnf(c_0_53, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_54, negated_conjecture, (~r1_tarski(k2_tarski(esk1_0,esk2_0),esk3_0)|~r2_hidden(esk1_0,esk3_0)|~r2_hidden(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_55, negated_conjecture, (r2_hidden(esk1_0,esk3_0)|r1_tarski(X1,esk3_0)|~r1_tarski(X1,k1_enumset1(esk1_0,esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_56, plain, (r1_tarski(k1_tarski(X1),k1_enumset1(X1,X1,X2))), inference(rw,[status(thm)],[c_0_44, c_0_21])).
cnf(c_0_57, plain, (esk5_2(k1_enumset1(X1,X1,X2),X3)=X1|r1_tarski(k1_enumset1(X1,X1,X2),X3)|~r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_58, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r1_tarski(k1_tarski(X1),X3)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_59, negated_conjecture, (r1_tarski(k1_tarski(esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_26, c_0_49])).
cnf(c_0_60, plain, (k2_xboole_0(X1,X2)=X1|~r1_tarski(k2_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_61, plain, (r1_tarski(k2_xboole_0(X1,X3),X2)|~r1_tarski(X1,X2)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_62, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_53])).
cnf(c_0_63, negated_conjecture, (~r2_hidden(esk1_0,esk3_0)|~r2_hidden(esk2_0,esk3_0)|~r1_tarski(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0)), inference(rw,[status(thm)],[c_0_54, c_0_21])).
cnf(c_0_64, negated_conjecture, (r2_hidden(esk1_0,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_47])).
cnf(c_0_65, plain, (r1_tarski(k1_enumset1(X1,X1,X2),X3)|~r2_hidden(X1,X3)|~r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_45, c_0_57])).
cnf(c_0_66, negated_conjecture, (r2_hidden(esk2_0,k2_xboole_0(X1,esk3_0))), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_67, plain, (k2_xboole_0(X1,X2)=X1|~r1_tarski(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62])])).
cnf(c_0_68, negated_conjecture, (~r2_hidden(esk2_0,esk3_0)|~r1_tarski(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_64])])).
cnf(c_0_69, negated_conjecture, (r1_tarski(k1_enumset1(esk1_0,esk1_0,X1),esk3_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_65, c_0_64])).
cnf(c_0_70, negated_conjecture, (r2_hidden(esk2_0,X1)|~r1_tarski(esk3_0,X1)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_71, negated_conjecture, (~r1_tarski(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_49])])).
cnf(c_0_72, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_62])]), c_0_71]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 73
# Proof object clause steps            : 46
# Proof object formula steps           : 27
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 26
# Processed clauses                    : 39640
# ...of these trivial                  : 526
# ...subsumed                          : 35072
# ...remaining for further processing  : 4042
# Other redundant clauses eliminated   : 117
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 124
# Backward-rewritten                   : 2049
# Generated clauses                    : 754403
# ...of the previous two non-trivial   : 674887
# Contextual simplify-reflections      : 18
# Paramodulations                      : 754224
# Factorizations                       : 58
# Equation resolutions                 : 121
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1839
#    Positive orientable unit clauses  : 300
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 1536
# Current number of unprocessed clauses: 632652
# ...number of literals in the above   : 2346551
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2199
# Clause-clause subsumption calls (NU) : 3397381
# Rec. Clause-clause subsumption calls : 1917180
# Non-unit clause-clause subsumptions  : 35203
# Unit Clause-clause subsumption calls : 6063
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 22007
# BW rewrite match successes           : 1606
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 10107135

# -------------------------------------------------
# User time                : 7.194 s
# System time              : 0.120 s
# Total time               : 7.314 s
# Maximum resident set size: 392404 pages
