# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t47_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2, X3]:((r1_yellow_0(X1,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)<=>r2_lattice3(X1,X3,X4))))=>k1_yellow_0(X1,X2)=k1_yellow_0(X1,X3))), file('data/MPTP/proofs/31b04d4cbcc548d4a3e0fd357023bfb2.E_input', t47_yellow_0)).
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/31b04d4cbcc548d4a3e0fd357023bfb2.E_input', d9_yellow_0)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/31b04d4cbcc548d4a3e0fd357023bfb2.E_input', dt_k1_yellow_0)).
fof(t46_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2, X3]:((![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)<=>r2_lattice3(X1,X3,X4)))&r1_yellow_0(X1,X2))=>r1_yellow_0(X1,X3))), file('data/MPTP/proofs/31b04d4cbcc548d4a3e0fd357023bfb2.E_input', t46_yellow_0)).
fof(c_0_4, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2, X3]:((r1_yellow_0(X1,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)<=>r2_lattice3(X1,X3,X4))))=>k1_yellow_0(X1,X2)=k1_yellow_0(X1,X3)))), inference(assume_negation,[status(cth)],[t47_yellow_0])).
fof(c_0_5, negated_conjecture, ![X8]:((~v2_struct_0(esk1_0)&l1_orders_2(esk1_0))&((r1_yellow_0(esk1_0,esk2_0)&((~r2_lattice3(esk1_0,esk2_0,X8)|r2_lattice3(esk1_0,esk3_0,X8)|~m1_subset_1(X8,u1_struct_0(esk1_0)))&(~r2_lattice3(esk1_0,esk3_0,X8)|r2_lattice3(esk1_0,esk2_0,X8)|~m1_subset_1(X8,u1_struct_0(esk1_0)))))&k1_yellow_0(esk1_0,esk2_0)!=k1_yellow_0(esk1_0,esk3_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_4])])])])])])).
fof(c_0_6, plain, ![X15, X16, X17, X18]:(((r2_lattice3(X15,X16,X17)|X17!=k1_yellow_0(X15,X16)|~r1_yellow_0(X15,X16)|~m1_subset_1(X17,u1_struct_0(X15))|~l1_orders_2(X15))&(~m1_subset_1(X18,u1_struct_0(X15))|(~r2_lattice3(X15,X16,X18)|r1_orders_2(X15,X17,X18))|X17!=k1_yellow_0(X15,X16)|~r1_yellow_0(X15,X16)|~m1_subset_1(X17,u1_struct_0(X15))|~l1_orders_2(X15)))&((m1_subset_1(esk5_3(X15,X16,X17),u1_struct_0(X15))|~r2_lattice3(X15,X16,X17)|X17=k1_yellow_0(X15,X16)|~r1_yellow_0(X15,X16)|~m1_subset_1(X17,u1_struct_0(X15))|~l1_orders_2(X15))&((r2_lattice3(X15,X16,esk5_3(X15,X16,X17))|~r2_lattice3(X15,X16,X17)|X17=k1_yellow_0(X15,X16)|~r1_yellow_0(X15,X16)|~m1_subset_1(X17,u1_struct_0(X15))|~l1_orders_2(X15))&(~r1_orders_2(X15,X17,esk5_3(X15,X16,X17))|~r2_lattice3(X15,X16,X17)|X17=k1_yellow_0(X15,X16)|~r1_yellow_0(X15,X16)|~m1_subset_1(X17,u1_struct_0(X15))|~l1_orders_2(X15))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
fof(c_0_7, plain, ![X9, X10]:(~l1_orders_2(X9)|m1_subset_1(k1_yellow_0(X9,X10),u1_struct_0(X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
cnf(c_0_8, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,X1)|~r2_lattice3(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(X1))|X3=k1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11, plain, (r2_lattice3(X1,X2,esk5_3(X1,X2,X3))|X3=k1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, plain, (r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~r1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14, negated_conjecture, (X1=k1_yellow_0(esk1_0,X2)|r2_lattice3(esk1_0,esk2_0,esk5_3(esk1_0,X2,X1))|~r2_lattice3(esk1_0,esk3_0,esk5_3(esk1_0,X2,X1))|~r2_lattice3(esk1_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10])])).
cnf(c_0_15, plain, (k1_yellow_0(X1,X2)=k1_yellow_0(X1,X3)|r2_lattice3(X1,X3,esk5_3(X1,X3,k1_yellow_0(X1,X2)))|~r2_lattice3(X1,X3,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X3)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),X3)|k1_yellow_0(X1,X2)!=k1_yellow_0(X1,X4)|~r2_lattice3(X1,X4,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_yellow_0(X1,X4)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_13, c_0_12])).
cnf(c_0_17, negated_conjecture, (k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|r2_lattice3(esk1_0,esk2_0,esk5_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_10])])).
cnf(c_0_18, plain, (X1=k1_yellow_0(X2,X3)|r1_orders_2(X2,k1_yellow_0(X2,X4),esk5_3(X2,X3,X1))|k1_yellow_0(X2,X4)!=k1_yellow_0(X2,X5)|~r2_lattice3(X2,X5,esk5_3(X2,X3,X1))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_yellow_0(X2,X5)|~r1_yellow_0(X2,X3)|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_16, c_0_9])).
cnf(c_0_19, negated_conjecture, (k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|r2_lattice3(esk1_0,esk2_0,esk5_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_12]), c_0_10])])).
cnf(c_0_20, negated_conjecture, (r1_yellow_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_21, negated_conjecture, (k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X2),esk5_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|k1_yellow_0(esk1_0,X2)!=k1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20]), c_0_10])])).
cnf(c_0_22, plain, (X2=k1_yellow_0(X1,X3)|~r1_orders_2(X1,X2,esk5_3(X1,X3,X2))|~r2_lattice3(X1,X3,X2)|~r1_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_23, negated_conjecture, (k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X2),esk5_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|k1_yellow_0(esk1_0,X2)!=k1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_12]), c_0_10])])).
cnf(c_0_24, negated_conjecture, (k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_10])])).
cnf(c_0_25, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,X1)|~r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_26, negated_conjecture, (k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_12]), c_0_10])])).
cnf(c_0_27, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_12]), c_0_10])])).
cnf(c_0_28, plain, (r2_lattice3(X1,X2,X3)|X3!=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_29, plain, ![X11, X12, X13]:((m1_subset_1(esk4_3(X11,X12,X13),u1_struct_0(X11))|~r1_yellow_0(X11,X12)|r1_yellow_0(X11,X13)|(v2_struct_0(X11)|~l1_orders_2(X11)))&((~r2_lattice3(X11,X12,esk4_3(X11,X12,X13))|~r2_lattice3(X11,X13,esk4_3(X11,X12,X13))|~r1_yellow_0(X11,X12)|r1_yellow_0(X11,X13)|(v2_struct_0(X11)|~l1_orders_2(X11)))&(r2_lattice3(X11,X12,esk4_3(X11,X12,X13))|r2_lattice3(X11,X13,esk4_3(X11,X12,X13))|~r1_yellow_0(X11,X12)|r1_yellow_0(X11,X13)|(v2_struct_0(X11)|~l1_orders_2(X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t46_yellow_0])])])])])])).
cnf(c_0_30, negated_conjecture, (k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_31, plain, (r2_lattice3(X1,X2,k1_yellow_0(X1,X3))|k1_yellow_0(X1,X3)!=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_28, c_0_12])).
cnf(c_0_32, plain, (r2_lattice3(X1,X2,esk4_3(X1,X2,X3))|r2_lattice3(X1,X3,esk4_3(X1,X2,X3))|r1_yellow_0(X1,X3)|v2_struct_0(X1)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_33, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_34, negated_conjecture, (k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,esk2_0)|~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_20]), c_0_10])])).
cnf(c_0_35, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,esk2_0,X1))|r2_lattice3(esk1_0,X1,esk4_3(esk1_0,esk2_0,X1))|r1_yellow_0(esk1_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_20]), c_0_10])]), c_0_33])).
cnf(c_0_36, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X3)|v2_struct_0(X1)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, negated_conjecture, (k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|r2_lattice3(esk1_0,esk3_0,esk4_3(esk1_0,esk2_0,esk3_0))|r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,esk2_0,esk3_0))|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_38, negated_conjecture, (k1_yellow_0(esk1_0,esk2_0)!=k1_yellow_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_39, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,X1,X2))|r1_yellow_0(esk1_0,X2)|~r2_lattice3(esk1_0,esk3_0,esk4_3(esk1_0,X1,X2))|~r1_yellow_0(esk1_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_36]), c_0_10])]), c_0_33])).
cnf(c_0_40, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,esk2_0,esk3_0))|r2_lattice3(esk1_0,esk3_0,esk4_3(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_37]), c_0_38])).
cnf(c_0_41, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,esk2_0,esk3_0))|r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_20])])).
cnf(c_0_42, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,esk4_3(esk1_0,X1,X2))|r1_yellow_0(esk1_0,X2)|~r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,X1,X2))|~r1_yellow_0(esk1_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_36]), c_0_10])]), c_0_33])).
cnf(c_0_43, negated_conjecture, (k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,esk2_0,esk3_0))|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_34, c_0_41])).
cnf(c_0_44, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,esk4_3(esk1_0,esk2_0,X1))|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,esk2_0,X1))), inference(spm,[status(thm)],[c_0_42, c_0_20])).
cnf(c_0_45, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_43]), c_0_38])).
cnf(c_0_46, negated_conjecture, (k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|r2_lattice3(esk1_0,esk3_0,esk4_3(esk1_0,esk2_0,esk3_0))|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_44]), c_0_45])])).
cnf(c_0_47, plain, (r1_yellow_0(X1,X3)|v2_struct_0(X1)|~r2_lattice3(X1,X2,esk4_3(X1,X2,X3))|~r2_lattice3(X1,X3,esk4_3(X1,X2,X3))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_48, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,esk4_3(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_46]), c_0_38])).
cnf(c_0_49, negated_conjecture, (r1_yellow_0(esk1_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_45]), c_0_20]), c_0_10])]), c_0_33])).
cnf(c_0_50, negated_conjecture, (k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk1_0,esk3_0)|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_49])])).
cnf(c_0_51, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_50]), c_0_38]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 52
# Proof object clause steps            : 43
# Proof object formula steps           : 9
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 44
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 15
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 15
# Processed clauses                    : 211
# ...of these trivial                  : 0
# ...subsumed                          : 35
# ...remaining for further processing  : 176
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 15
# Backward-rewritten                   : 37
# Generated clauses                    : 675
# ...of the previous two non-trivial   : 617
# Contextual simplify-reflections      : 7
# Paramodulations                      : 614
# Factorizations                       : 38
# Equation resolutions                 : 23
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 109
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 102
# Current number of unprocessed clauses: 429
# ...number of literals in the above   : 3762
# Current number of archived formulas  : 0
# Current number of archived clauses   : 67
# Clause-clause subsumption calls (NU) : 5101
# Rec. Clause-clause subsumption calls : 794
# Non-unit clause-clause subsumptions  : 55
# Unit Clause-clause subsumption calls : 217
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 39
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 27773

# -------------------------------------------------
# User time                : 0.030 s
# System time              : 0.000 s
# Total time               : 0.030 s
# Maximum resident set size: 28608 pages
