# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t122_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_pre_topc(X2,X1)=>((v1_tsep_1(X2,X1)&m1_pre_topc(X2,X1))<=>(v1_funct_1(k9_tmap_1(X1,X2))&(v1_funct_2(k9_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))&(v5_pre_topc(k9_tmap_1(X1,X2),X1,k8_tmap_1(X1,X2))&m1_subset_1(k9_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2))))))))))), file('data/MPTP/proofs/320d9b541c274998b227b464d62363c8.E_input', t122_tmap_1)).
fof(d11_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:((v1_pre_topc(X3)&(v2_pre_topc(X3)&l1_pre_topc(X3)))=>(X3=k8_tmap_1(X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(X4=u1_struct_0(X2)=>X3=k6_tmap_1(X1,X4))))))), file('data/MPTP/proofs/320d9b541c274998b227b464d62363c8.E_input', d11_tmap_1)).
fof(fc4_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(~(v2_struct_0(k6_tmap_1(X1,X2)))&(v1_pre_topc(k6_tmap_1(X1,X2))&v2_pre_topc(k6_tmap_1(X1,X2))))), file('data/MPTP/proofs/320d9b541c274998b227b464d62363c8.E_input', fc4_tmap_1)).
fof(t1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/320d9b541c274998b227b464d62363c8.E_input', t1_tsep_1)).
fof(dt_k6_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_pre_topc(k6_tmap_1(X1,X2))&(v2_pre_topc(k6_tmap_1(X1,X2))&l1_pre_topc(k6_tmap_1(X1,X2))))), file('data/MPTP/proofs/320d9b541c274998b227b464d62363c8.E_input', dt_k6_tmap_1)).
fof(t104_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(u1_struct_0(k6_tmap_1(X1,X2))=u1_struct_0(X1)&u1_pre_topc(k6_tmap_1(X1,X2))=k5_tmap_1(X1,X2)))), file('data/MPTP/proofs/320d9b541c274998b227b464d62363c8.E_input', t104_tmap_1)).
fof(t2_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>m1_pre_topc(X1,X1)), file('data/MPTP/proofs/320d9b541c274998b227b464d62363c8.E_input', t2_tsep_1)).
fof(t39_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/320d9b541c274998b227b464d62363c8.E_input', t39_pre_topc)).
fof(d1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>(v1_tsep_1(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>v3_pre_topc(X3,X1)))))), file('data/MPTP/proofs/320d9b541c274998b227b464d62363c8.E_input', d1_tsep_1)).
fof(t113_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>(v1_funct_1(k7_tmap_1(X1,X2))&(v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))&(v5_pre_topc(k7_tmap_1(X1,X2),X1,k6_tmap_1(X1,X2))&m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2))))))))))), file('data/MPTP/proofs/320d9b541c274998b227b464d62363c8.E_input', t113_tmap_1)).
fof(dt_k7_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_funct_1(k7_tmap_1(X1,X2))&(v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))&m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))))), file('data/MPTP/proofs/320d9b541c274998b227b464d62363c8.E_input', dt_k7_tmap_1)).
fof(d10_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k7_tmap_1(X1,X2)=k6_partfun1(u1_struct_0(X1)))), file('data/MPTP/proofs/320d9b541c274998b227b464d62363c8.E_input', d10_tmap_1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/320d9b541c274998b227b464d62363c8.E_input', fc2_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/320d9b541c274998b227b464d62363c8.E_input', dt_l1_pre_topc)).
fof(dt_k9_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_pre_topc(X2,X1))=>(v1_funct_1(k9_tmap_1(X1,X2))&(v1_funct_2(k9_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))&m1_subset_1(k9_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))))), file('data/MPTP/proofs/320d9b541c274998b227b464d62363c8.E_input', dt_k9_tmap_1)).
fof(d12_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))))=>(X3=k9_tmap_1(X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(X4=u1_struct_0(X2)=>r1_funct_2(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X4)),X3,k7_tmap_1(X1,X4)))))))), file('data/MPTP/proofs/320d9b541c274998b227b464d62363c8.E_input', d12_tmap_1)).
fof(reflexivity_r1_funct_2, axiom, ![X1, X2, X3, X4, X5, X6]:((~(v1_xboole_0(X2))&(~(v1_xboole_0(X4))&((v1_funct_1(X5)&(v1_funct_2(X5,X1,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&(v1_funct_1(X6)&(v1_funct_2(X6,X3,X4)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))))))=>r1_funct_2(X1,X2,X3,X4,X5,X5)), file('data/MPTP/proofs/320d9b541c274998b227b464d62363c8.E_input', reflexivity_r1_funct_2)).
fof(c_0_17, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_pre_topc(X2,X1)=>((v1_tsep_1(X2,X1)&m1_pre_topc(X2,X1))<=>(v1_funct_1(k9_tmap_1(X1,X2))&(v1_funct_2(k9_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))&(v5_pre_topc(k9_tmap_1(X1,X2),X1,k8_tmap_1(X1,X2))&m1_subset_1(k9_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))))))))), inference(assume_negation,[status(cth)],[t122_tmap_1])).
fof(c_0_18, plain, ![X16, X17, X18, X19]:((X18!=k8_tmap_1(X16,X17)|(~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X16)))|(X19!=u1_struct_0(X17)|X18=k6_tmap_1(X16,X19)))|(~v1_pre_topc(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18)))|~m1_pre_topc(X17,X16)|(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))))&((m1_subset_1(esk3_3(X16,X17,X18),k1_zfmisc_1(u1_struct_0(X16)))|X18=k8_tmap_1(X16,X17)|(~v1_pre_topc(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18)))|~m1_pre_topc(X17,X16)|(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))))&((esk3_3(X16,X17,X18)=u1_struct_0(X17)|X18=k8_tmap_1(X16,X17)|(~v1_pre_topc(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18)))|~m1_pre_topc(X17,X16)|(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))))&(X18!=k6_tmap_1(X16,esk3_3(X16,X17,X18))|X18=k8_tmap_1(X16,X17)|(~v1_pre_topc(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18)))|~m1_pre_topc(X17,X16)|(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d11_tmap_1])])])])])])).
fof(c_0_19, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_pre_topc(esk2_0,esk1_0)&((~v1_tsep_1(esk2_0,esk1_0)|~m1_pre_topc(esk2_0,esk1_0)|(~v1_funct_1(k9_tmap_1(esk1_0,esk2_0))|(~v1_funct_2(k9_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))|(~v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))|~m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))))))&(((v1_funct_1(k9_tmap_1(esk1_0,esk2_0))|v1_tsep_1(esk2_0,esk1_0))&((v1_funct_2(k9_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))|v1_tsep_1(esk2_0,esk1_0))&((v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))|v1_tsep_1(esk2_0,esk1_0))&(m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))|v1_tsep_1(esk2_0,esk1_0)))))&((v1_funct_1(k9_tmap_1(esk1_0,esk2_0))|m1_pre_topc(esk2_0,esk1_0))&((v1_funct_2(k9_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))|m1_pre_topc(esk2_0,esk1_0))&((v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))|m1_pre_topc(esk2_0,esk1_0))&(m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))|m1_pre_topc(esk2_0,esk1_0))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])])).
fof(c_0_20, plain, ![X21, X22]:((~v2_struct_0(k6_tmap_1(X21,X22))|(v2_struct_0(X21)|(~v2_pre_topc(X21)|~l1_pre_topc(X21))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))))&((v1_pre_topc(k6_tmap_1(X21,X22))|(v2_struct_0(X21)|(~v2_pre_topc(X21)|~l1_pre_topc(X21))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))))&(v2_pre_topc(k6_tmap_1(X21,X22))|(v2_struct_0(X21)|(~v2_pre_topc(X21)|~l1_pre_topc(X21))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc4_tmap_1])])])])).
fof(c_0_21, plain, ![X12, X13]:(~l1_pre_topc(X12)|(~m1_pre_topc(X13,X12)|m1_subset_1(u1_struct_0(X13),k1_zfmisc_1(u1_struct_0(X12))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])).
fof(c_0_22, plain, ![X23, X24]:((v1_pre_topc(k6_tmap_1(X23,X24))|(v2_struct_0(X23)|(~v2_pre_topc(X23)|~l1_pre_topc(X23))|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X23)))))&((v2_pre_topc(k6_tmap_1(X23,X24))|(v2_struct_0(X23)|(~v2_pre_topc(X23)|~l1_pre_topc(X23))|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X23)))))&(l1_pre_topc(k6_tmap_1(X23,X24))|(v2_struct_0(X23)|(~v2_pre_topc(X23)|~l1_pre_topc(X23))|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X23))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_tmap_1])])])])).
cnf(c_0_23, plain, (esk3_3(X1,X2,X3)=u1_struct_0(X2)|X3=k8_tmap_1(X1,X2)|v2_struct_0(X1)|~v1_pre_topc(X3)|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, plain, (v1_pre_topc(k6_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, plain, (l1_pre_topc(k6_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, plain, (v2_pre_topc(k6_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_32, plain, ![X25, X26]:((u1_struct_0(k6_tmap_1(X25,X26))=u1_struct_0(X25)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|(v2_struct_0(X25)|(~v2_pre_topc(X25)|~l1_pre_topc(X25))))&(u1_pre_topc(k6_tmap_1(X25,X26))=k5_tmap_1(X25,X26)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|(v2_struct_0(X25)|(~v2_pre_topc(X25)|~l1_pre_topc(X25))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t104_tmap_1])])])])])).
cnf(c_0_33, plain, (X1=k8_tmap_1(X2,X3)|v2_struct_0(X2)|X1!=k6_tmap_1(X2,esk3_3(X2,X3,X1))|~v1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_34, negated_conjecture, (esk3_3(esk1_0,esk2_0,X1)=u1_struct_0(esk2_0)|X1=k8_tmap_1(esk1_0,esk2_0)|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_35, plain, (v1_pre_topc(k6_tmap_1(X1,u1_struct_0(X2)))|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_36, plain, (l1_pre_topc(k6_tmap_1(X1,u1_struct_0(X2)))|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_30, c_0_29])).
cnf(c_0_37, plain, (v2_pre_topc(k6_tmap_1(X1,u1_struct_0(X2)))|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_31, c_0_29])).
cnf(c_0_38, plain, (u1_struct_0(k6_tmap_1(X1,X2))=u1_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39, negated_conjecture, (X1=k8_tmap_1(esk1_0,esk2_0)|k6_tmap_1(esk1_0,u1_struct_0(esk2_0))!=X1|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_40, negated_conjecture, (v1_pre_topc(k6_tmap_1(esk1_0,u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_41, negated_conjecture, (l1_pre_topc(k6_tmap_1(esk1_0,u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_42, negated_conjecture, (v2_pre_topc(k6_tmap_1(esk1_0,u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_43, plain, (u1_struct_0(k6_tmap_1(X1,u1_struct_0(X2)))=u1_struct_0(X1)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_38, c_0_29])).
cnf(c_0_44, negated_conjecture, (k6_tmap_1(esk1_0,u1_struct_0(esk2_0))=k8_tmap_1(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41]), c_0_42])])).
cnf(c_0_45, negated_conjecture, (u1_struct_0(k8_tmap_1(esk1_0,esk2_0))=u1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_24]), c_0_44]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_46, negated_conjecture, (l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_41, c_0_44])).
fof(c_0_47, plain, ![X27]:(~l1_pre_topc(X27)|m1_pre_topc(X27,X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tsep_1])])).
cnf(c_0_48, negated_conjecture, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X1,k8_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_45]), c_0_46])])).
cnf(c_0_49, plain, (m1_pre_topc(X1,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_50, plain, ![X9, X10, X11]:(~l1_pre_topc(X9)|(~m1_pre_topc(X10,X9)|(~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X9)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t39_pre_topc])])])).
cnf(c_0_51, negated_conjecture, (m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_45]), c_0_46])])).
cnf(c_0_52, plain, (m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_53, negated_conjecture, (u1_struct_0(k6_tmap_1(esk1_0,u1_struct_0(esk1_0)))=u1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_51]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(k6_tmap_1(esk1_0,u1_struct_0(esk1_0)),X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
fof(c_0_55, plain, ![X49, X50, X51]:((~v1_tsep_1(X50,X49)|(~m1_subset_1(X51,k1_zfmisc_1(u1_struct_0(X49)))|(X51!=u1_struct_0(X50)|v3_pre_topc(X51,X49)))|~m1_pre_topc(X50,X49)|~l1_pre_topc(X49))&((m1_subset_1(esk5_2(X49,X50),k1_zfmisc_1(u1_struct_0(X49)))|v1_tsep_1(X50,X49)|~m1_pre_topc(X50,X49)|~l1_pre_topc(X49))&((esk5_2(X49,X50)=u1_struct_0(X50)|v1_tsep_1(X50,X49)|~m1_pre_topc(X50,X49)|~l1_pre_topc(X49))&(~v3_pre_topc(esk5_2(X49,X50),X49)|v1_tsep_1(X50,X49)|~m1_pre_topc(X50,X49)|~l1_pre_topc(X49))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tsep_1])])])])])).
cnf(c_0_56, negated_conjecture, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(k6_tmap_1(esk1_0,u1_struct_0(esk1_0)),X2)|~m1_pre_topc(X1,esk1_0)|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_29]), c_0_25])])).
cnf(c_0_57, plain, (m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|v1_tsep_1(X2,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_58, plain, (v2_struct_0(X1)|~v2_struct_0(k6_tmap_1(X1,X2))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_59, negated_conjecture, (m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_pre_topc(k6_tmap_1(esk1_0,u1_struct_0(esk1_0)),X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_56, c_0_24])).
cnf(c_0_60, negated_conjecture, (l1_pre_topc(k6_tmap_1(esk1_0,u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_51]), c_0_25]), c_0_26])]), c_0_27])).
fof(c_0_61, plain, ![X47, X48]:(((v1_funct_1(k7_tmap_1(X47,X48))|~v3_pre_topc(X48,X47)|~m1_subset_1(X48,k1_zfmisc_1(u1_struct_0(X47)))|(v2_struct_0(X47)|(~v2_pre_topc(X47)|~l1_pre_topc(X47))))&((v1_funct_2(k7_tmap_1(X47,X48),u1_struct_0(X47),u1_struct_0(k6_tmap_1(X47,X48)))|~v3_pre_topc(X48,X47)|~m1_subset_1(X48,k1_zfmisc_1(u1_struct_0(X47)))|(v2_struct_0(X47)|(~v2_pre_topc(X47)|~l1_pre_topc(X47))))&((v5_pre_topc(k7_tmap_1(X47,X48),X47,k6_tmap_1(X47,X48))|~v3_pre_topc(X48,X47)|~m1_subset_1(X48,k1_zfmisc_1(u1_struct_0(X47)))|(v2_struct_0(X47)|(~v2_pre_topc(X47)|~l1_pre_topc(X47))))&(m1_subset_1(k7_tmap_1(X47,X48),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X47),u1_struct_0(k6_tmap_1(X47,X48)))))|~v3_pre_topc(X48,X47)|~m1_subset_1(X48,k1_zfmisc_1(u1_struct_0(X47)))|(v2_struct_0(X47)|(~v2_pre_topc(X47)|~l1_pre_topc(X47)))))))&(~v1_funct_1(k7_tmap_1(X47,X48))|(~v1_funct_2(k7_tmap_1(X47,X48),u1_struct_0(X47),u1_struct_0(k6_tmap_1(X47,X48)))|(~v5_pre_topc(k7_tmap_1(X47,X48),X47,k6_tmap_1(X47,X48))|~m1_subset_1(k7_tmap_1(X47,X48),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X47),u1_struct_0(k6_tmap_1(X47,X48)))))))|v3_pre_topc(X48,X47)|~m1_subset_1(X48,k1_zfmisc_1(u1_struct_0(X47)))|(v2_struct_0(X47)|(~v2_pre_topc(X47)|~l1_pre_topc(X47))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t113_tmap_1])])])])])).
fof(c_0_62, plain, ![X39, X40]:((v1_funct_1(k7_tmap_1(X39,X40))|(v2_struct_0(X39)|(~v2_pre_topc(X39)|~l1_pre_topc(X39))|~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X39)))))&((v1_funct_2(k7_tmap_1(X39,X40),u1_struct_0(X39),u1_struct_0(k6_tmap_1(X39,X40)))|(v2_struct_0(X39)|(~v2_pre_topc(X39)|~l1_pre_topc(X39))|~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X39)))))&(m1_subset_1(k7_tmap_1(X39,X40),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X39),u1_struct_0(k6_tmap_1(X39,X40)))))|(v2_struct_0(X39)|(~v2_pre_topc(X39)|~l1_pre_topc(X39))|~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X39))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_tmap_1])])])])).
cnf(c_0_63, negated_conjecture, (m1_subset_1(esk5_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_24]), c_0_25])])).
fof(c_0_64, plain, ![X28, X29]:(v2_struct_0(X28)|(~v2_pre_topc(X28)|~l1_pre_topc(X28))|(~m1_subset_1(X29,k1_zfmisc_1(u1_struct_0(X28)))|k7_tmap_1(X28,X29)=k6_partfun1(u1_struct_0(X28)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_tmap_1])])])])).
fof(c_0_65, plain, ![X30]:(v2_struct_0(X30)|~l1_struct_0(X30)|~v1_xboole_0(u1_struct_0(X30))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_66, plain, ![X31]:(~l1_pre_topc(X31)|l1_struct_0(X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_67, negated_conjecture, (~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_struct_0(k8_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_44]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_68, negated_conjecture, (m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_49]), c_0_53]), c_0_60])])).
cnf(c_0_69, plain, (v3_pre_topc(X2,X1)|v2_struct_0(X1)|~v1_funct_1(k7_tmap_1(X1,X2))|~v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))|~v5_pre_topc(k7_tmap_1(X1,X2),X1,k6_tmap_1(X1,X2))|~m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_70, plain, (v1_funct_1(k7_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_71, plain, (v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_72, plain, (m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_73, negated_conjecture, (u1_struct_0(k6_tmap_1(esk1_0,esk5_2(esk1_0,esk2_0)))=u1_struct_0(esk1_0)|v1_tsep_1(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_63]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_74, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|l1_pre_topc(k6_tmap_1(esk1_0,esk5_2(esk1_0,esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_63]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_75, plain, (v2_struct_0(X1)|k7_tmap_1(X1,X2)=k6_partfun1(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_64])).
fof(c_0_76, plain, ![X32, X33]:((v1_funct_1(k9_tmap_1(X32,X33))|(v2_struct_0(X32)|(~v2_pre_topc(X32)|~l1_pre_topc(X32))|~m1_pre_topc(X33,X32)))&((v1_funct_2(k9_tmap_1(X32,X33),u1_struct_0(X32),u1_struct_0(k8_tmap_1(X32,X33)))|(v2_struct_0(X32)|(~v2_pre_topc(X32)|~l1_pre_topc(X32))|~m1_pre_topc(X33,X32)))&(m1_subset_1(k9_tmap_1(X32,X33),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X32),u1_struct_0(k8_tmap_1(X32,X33)))))|(v2_struct_0(X32)|(~v2_pre_topc(X32)|~l1_pre_topc(X32))|~m1_pre_topc(X33,X32))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k9_tmap_1])])])])).
cnf(c_0_77, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_78, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_79, plain, (u1_struct_0(k6_tmap_1(X1,u1_struct_0(X1)))=u1_struct_0(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_43, c_0_49])).
cnf(c_0_80, negated_conjecture, (v2_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_42, c_0_44])).
cnf(c_0_81, negated_conjecture, (~v2_struct_0(k8_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_68])])).
cnf(c_0_82, plain, (l1_pre_topc(k6_tmap_1(X1,u1_struct_0(X1)))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_36, c_0_49])).
cnf(c_0_83, plain, (v3_pre_topc(X1,X2)|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v5_pre_topc(k7_tmap_1(X2,X1),X2,k6_tmap_1(X2,X1))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_69, c_0_70]), c_0_71]), c_0_72])).
cnf(c_0_84, plain, (esk5_2(X1,X2)=u1_struct_0(X2)|v1_tsep_1(X2,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_85, negated_conjecture, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|v1_tsep_1(esk2_0,esk1_0)|~m1_pre_topc(X1,k6_tmap_1(esk1_0,esk5_2(esk1_0,esk2_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_73]), c_0_74])).
fof(c_0_86, plain, ![X34, X35, X36, X37]:((X36!=k9_tmap_1(X34,X35)|(~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X34)))|(X37!=u1_struct_0(X35)|r1_funct_2(u1_struct_0(X34),u1_struct_0(k8_tmap_1(X34,X35)),u1_struct_0(X34),u1_struct_0(k6_tmap_1(X34,X37)),X36,k7_tmap_1(X34,X37))))|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(X34),u1_struct_0(k8_tmap_1(X34,X35)))|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X34),u1_struct_0(k8_tmap_1(X34,X35)))))))|~m1_pre_topc(X35,X34)|(v2_struct_0(X34)|(~v2_pre_topc(X34)|~l1_pre_topc(X34))))&((m1_subset_1(esk4_3(X34,X35,X36),k1_zfmisc_1(u1_struct_0(X34)))|X36=k9_tmap_1(X34,X35)|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(X34),u1_struct_0(k8_tmap_1(X34,X35)))|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X34),u1_struct_0(k8_tmap_1(X34,X35)))))))|~m1_pre_topc(X35,X34)|(v2_struct_0(X34)|(~v2_pre_topc(X34)|~l1_pre_topc(X34))))&((esk4_3(X34,X35,X36)=u1_struct_0(X35)|X36=k9_tmap_1(X34,X35)|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(X34),u1_struct_0(k8_tmap_1(X34,X35)))|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X34),u1_struct_0(k8_tmap_1(X34,X35)))))))|~m1_pre_topc(X35,X34)|(v2_struct_0(X34)|(~v2_pre_topc(X34)|~l1_pre_topc(X34))))&(~r1_funct_2(u1_struct_0(X34),u1_struct_0(k8_tmap_1(X34,X35)),u1_struct_0(X34),u1_struct_0(k6_tmap_1(X34,esk4_3(X34,X35,X36))),X36,k7_tmap_1(X34,esk4_3(X34,X35,X36)))|X36=k9_tmap_1(X34,X35)|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(X34),u1_struct_0(k8_tmap_1(X34,X35)))|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X34),u1_struct_0(k8_tmap_1(X34,X35)))))))|~m1_pre_topc(X35,X34)|(v2_struct_0(X34)|(~v2_pre_topc(X34)|~l1_pre_topc(X34))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d12_tmap_1])])])])])])).
cnf(c_0_87, negated_conjecture, (k6_partfun1(u1_struct_0(esk1_0))=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_51]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_88, plain, (m1_subset_1(k9_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_89, plain, (v1_funct_2(k9_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_90, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_91, negated_conjecture, (u1_struct_0(k6_tmap_1(k8_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0)))=u1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_45]), c_0_46]), c_0_80])]), c_0_81])).
cnf(c_0_92, negated_conjecture, (l1_pre_topc(k6_tmap_1(k8_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_45]), c_0_46]), c_0_80])]), c_0_81])).
cnf(c_0_93, negated_conjecture, (v3_pre_topc(esk5_2(esk1_0,esk2_0),esk1_0)|v1_tsep_1(esk2_0,esk1_0)|~v5_pre_topc(k7_tmap_1(esk1_0,esk5_2(esk1_0,esk2_0)),esk1_0,k6_tmap_1(esk1_0,esk5_2(esk1_0,esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_63]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_94, negated_conjecture, (k7_tmap_1(esk1_0,esk5_2(esk1_0,esk2_0))=k6_partfun1(u1_struct_0(esk1_0))|v1_tsep_1(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_63]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_95, plain, (v1_tsep_1(X2,X1)|~v3_pre_topc(esk5_2(X1,X2),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_96, negated_conjecture, (esk5_2(esk1_0,esk2_0)=u1_struct_0(esk2_0)|v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_24]), c_0_25])])).
cnf(c_0_97, negated_conjecture, (m1_subset_1(u1_struct_0(k6_tmap_1(esk1_0,esk5_2(esk1_0,esk2_0))),k1_zfmisc_1(u1_struct_0(esk1_0)))|v1_tsep_1(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_49]), c_0_74])).
cnf(c_0_98, plain, (esk4_3(X1,X2,X3)=u1_struct_0(X2)|X3=k9_tmap_1(X1,X2)|v2_struct_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_99, negated_conjecture, (k7_tmap_1(esk1_0,u1_struct_0(esk2_0))=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_68]), c_0_87]), c_0_25]), c_0_26])]), c_0_27])).
fof(c_0_100, plain, ![X41, X42, X43, X44, X45, X46]:(v1_xboole_0(X42)|(v1_xboole_0(X44)|(~v1_funct_1(X45)|(~v1_funct_2(X45,X41,X42)|~m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(X41,X42))))|(~v1_funct_1(X46)|(~v1_funct_2(X46,X43,X44)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(X43,X44)))))))|r1_funct_2(X41,X42,X43,X44,X45,X45)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r1_funct_2])])])).
cnf(c_0_101, negated_conjecture, (m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_102, negated_conjecture, (v1_funct_2(k9_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_103, plain, (v1_funct_1(k9_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_104, negated_conjecture, (v2_struct_0(k6_tmap_1(k8_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0)))|~v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_92])])).
cnf(c_0_105, negated_conjecture, (~v1_tsep_1(esk2_0,esk1_0)|~m1_pre_topc(esk2_0,esk1_0)|~v1_funct_1(k9_tmap_1(esk1_0,esk2_0))|~v1_funct_2(k9_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))|~v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))|~m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_106, negated_conjecture, (v3_pre_topc(esk5_2(esk1_0,esk2_0),esk1_0)|v1_tsep_1(esk2_0,esk1_0)|~v5_pre_topc(k6_partfun1(u1_struct_0(esk1_0)),esk1_0,k6_tmap_1(esk1_0,esk5_2(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_107, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|~v3_pre_topc(u1_struct_0(esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_24]), c_0_25])])).
cnf(c_0_108, negated_conjecture, (k7_tmap_1(esk1_0,u1_struct_0(k6_tmap_1(esk1_0,esk5_2(esk1_0,esk2_0))))=k6_partfun1(u1_struct_0(esk1_0))|v1_tsep_1(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_97]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_109, plain, (X3=k9_tmap_1(X1,X2)|v2_struct_0(X1)|~r1_funct_2(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,esk4_3(X1,X2,X3))),X3,k7_tmap_1(X1,esk4_3(X1,X2,X3)))|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_110, negated_conjecture, (esk4_3(esk1_0,esk2_0,X1)=u1_struct_0(esk2_0)|X1=k9_tmap_1(esk1_0,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_45]), c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_111, negated_conjecture, (v1_funct_2(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_68]), c_0_44]), c_0_45]), c_0_25]), c_0_26])]), c_0_27]), c_0_99])).
cnf(c_0_112, negated_conjecture, (m1_subset_1(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_68]), c_0_44]), c_0_45]), c_0_25]), c_0_26])]), c_0_27]), c_0_99])).
cnf(c_0_113, negated_conjecture, (v1_funct_1(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_51]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_114, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|r1_funct_2(X4,X1,X6,X2,X3,X3)|~v1_funct_1(X3)|~v1_funct_2(X3,X4,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X1)))|~v1_funct_1(X5)|~v1_funct_2(X5,X6,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X2)))), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_115, negated_conjecture, (m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(rw,[status(thm)],[c_0_101, c_0_45])).
cnf(c_0_116, negated_conjecture, (v1_funct_2(k9_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(rw,[status(thm)],[c_0_102, c_0_45])).
cnf(c_0_117, negated_conjecture, (v1_funct_1(k9_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_118, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_104]), c_0_45]), c_0_51]), c_0_46]), c_0_80])]), c_0_81])).
cnf(c_0_119, negated_conjecture, (~m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))|~v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))|~v1_funct_2(k9_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))|~v1_funct_1(k9_tmap_1(esk1_0,esk2_0))|~v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_105, c_0_24])])).
cnf(c_0_120, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|~v5_pre_topc(k6_partfun1(u1_struct_0(esk1_0)),esk1_0,k6_tmap_1(esk1_0,u1_struct_0(esk2_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_96]), c_0_107])).
cnf(c_0_121, negated_conjecture, (k6_partfun1(u1_struct_0(esk1_0))=k7_tmap_1(esk1_0,u1_struct_0(esk1_0))|v1_tsep_1(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_108, c_0_73])).
cnf(c_0_122, negated_conjecture, (X1=k9_tmap_1(esk1_0,esk2_0)|~r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,esk4_3(esk1_0,esk2_0,X1))),X1,k7_tmap_1(esk1_0,esk4_3(esk1_0,esk2_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_45]), c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_123, negated_conjecture, (esk4_3(esk1_0,esk2_0,k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))=u1_struct_0(esk2_0)|k7_tmap_1(esk1_0,u1_struct_0(esk1_0))=k9_tmap_1(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_112]), c_0_113])])).
cnf(c_0_124, negated_conjecture, (r1_funct_2(X1,X2,u1_struct_0(esk1_0),u1_struct_0(esk1_0),X3,X3)|v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_116]), c_0_117])]), c_0_118])).
cnf(c_0_125, negated_conjecture, (~m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))|~v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))|~v1_funct_2(k9_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))|~v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_119, c_0_117])])).
cnf(c_0_126, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|~v5_pre_topc(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),esk1_0,k6_tmap_1(esk1_0,u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_127, negated_conjecture, (k7_tmap_1(esk1_0,u1_struct_0(esk1_0))=k9_tmap_1(esk1_0,esk2_0)|~r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_44]), c_0_45]), c_0_99]), c_0_112]), c_0_111]), c_0_113])])).
cnf(c_0_128, negated_conjecture, (r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),k7_tmap_1(esk1_0,u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_112]), c_0_111]), c_0_113])]), c_0_118])).
cnf(c_0_129, negated_conjecture, (~m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))|~v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))|~v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_125, c_0_102])])).
cnf(c_0_130, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|~v5_pre_topc(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),esk1_0,k8_tmap_1(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_126, c_0_44])).
cnf(c_0_131, negated_conjecture, (k7_tmap_1(esk1_0,u1_struct_0(esk1_0))=k9_tmap_1(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_127, c_0_128])])).
cnf(c_0_132, negated_conjecture, (v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))|v1_tsep_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_133, plain, (v5_pre_topc(k7_tmap_1(X1,X2),X1,k6_tmap_1(X1,X2))|v2_struct_0(X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_134, negated_conjecture, (~v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))|~v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_129, c_0_101])])).
cnf(c_0_135, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[c_0_130, c_0_131]), c_0_132])).
cnf(c_0_136, negated_conjecture, (v5_pre_topc(k7_tmap_1(esk1_0,u1_struct_0(esk1_0)),esk1_0,k8_tmap_1(esk1_0,esk2_0))|~v3_pre_topc(u1_struct_0(esk2_0),esk1_0)), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_68]), c_0_44]), c_0_25]), c_0_26])]), c_0_27]), c_0_99])).
cnf(c_0_137, negated_conjecture, (~v5_pre_topc(k9_tmap_1(esk1_0,esk2_0),esk1_0,k8_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_134, c_0_135])])).
cnf(c_0_138, plain, (v3_pre_topc(X3,X2)|~v1_tsep_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|X3!=u1_struct_0(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_139, negated_conjecture, (~v3_pre_topc(u1_struct_0(esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_136, c_0_131]), c_0_137])).
cnf(c_0_140, plain, (v3_pre_topc(u1_struct_0(X1),X2)|~v1_tsep_1(X1,X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_138]), c_0_29])).
cnf(c_0_141, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140]), c_0_135]), c_0_24]), c_0_25])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 142
# Proof object clause steps            : 107
# Proof object formula steps           : 35
# Proof object conjectures             : 73
# Proof object clause conjectures      : 70
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 34
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 58
# Proof object simplifying inferences  : 201
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 52
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 52
# Processed clauses                    : 2026
# ...of these trivial                  : 17
# ...subsumed                          : 1031
# ...remaining for further processing  : 978
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 67
# Backward-rewritten                   : 376
# Generated clauses                    : 7202
# ...of the previous two non-trivial   : 5809
# Contextual simplify-reflections      : 212
# Paramodulations                      : 7177
# Factorizations                       : 0
# Equation resolutions                 : 15
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 482
#    Positive orientable unit clauses  : 32
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 445
# Current number of unprocessed clauses: 3569
# ...number of literals in the above   : 19308
# Current number of archived formulas  : 0
# Current number of archived clauses   : 496
# Clause-clause subsumption calls (NU) : 111836
# Rec. Clause-clause subsumption calls : 34184
# Non-unit clause-clause subsumptions  : 1287
# Unit Clause-clause subsumption calls : 2345
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 62
# BW rewrite match successes           : 29
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 255467

# -------------------------------------------------
# User time                : 0.150 s
# System time              : 0.007 s
# Total time               : 0.156 s
# Maximum resident set size: 28816 pages
