# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/3223abcad9d146c9a2641cc4d541619c.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/3223abcad9d146c9a2641cc4d541619c.E_input', t70_enumset1)).
fof(d3_xboole_0, axiom, ![X1, X2, X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2)))), file('data/MPTP/proofs/3223abcad9d146c9a2641cc4d541619c.E_input', d3_xboole_0)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/3223abcad9d146c9a2641cc4d541619c.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/3223abcad9d146c9a2641cc4d541619c.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/3223abcad9d146c9a2641cc4d541619c.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/3223abcad9d146c9a2641cc4d541619c.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/3223abcad9d146c9a2641cc4d541619c.E_input', t75_enumset1)).
fof(t34_ordinal1, conjecture, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,k1_ordinal1(X2))<=>r1_ordinal1(X1,X2)))), file('data/MPTP/proofs/3223abcad9d146c9a2641cc4d541619c.E_input', t34_ordinal1)).
fof(d1_ordinal1, axiom, ![X1]:k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1)), file('data/MPTP/proofs/3223abcad9d146c9a2641cc4d541619c.E_input', d1_ordinal1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/3223abcad9d146c9a2641cc4d541619c.E_input', t69_enumset1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/3223abcad9d146c9a2641cc4d541619c.E_input', d1_tarski)).
fof(redefinition_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)<=>r1_tarski(X1,X2))), file('data/MPTP/proofs/3223abcad9d146c9a2641cc4d541619c.E_input', redefinition_r1_ordinal1)).
fof(d2_ordinal1, axiom, ![X1]:(v1_ordinal1(X1)<=>![X2]:(r2_hidden(X2,X1)=>r1_tarski(X2,X1))), file('data/MPTP/proofs/3223abcad9d146c9a2641cc4d541619c.E_input', d2_ordinal1)).
fof(cc1_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>(v1_ordinal1(X1)&v2_ordinal1(X1))), file('data/MPTP/proofs/3223abcad9d146c9a2641cc4d541619c.E_input', cc1_ordinal1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/3223abcad9d146c9a2641cc4d541619c.E_input', d10_xboole_0)).
fof(t26_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r1_ordinal1(X1,X2)|r2_hidden(X2,X1)))), file('data/MPTP/proofs/3223abcad9d146c9a2641cc4d541619c.E_input', t26_ordinal1)).
fof(connectedness_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)|r1_ordinal1(X2,X1))), file('data/MPTP/proofs/3223abcad9d146c9a2641cc4d541619c.E_input', connectedness_r1_ordinal1)).
fof(d6_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:(X9=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)<=>![X10]:(r2_hidden(X10,X9)<=>~((~(X10=X1)&(~(X10=X2)&(~(X10=X3)&(~(X10=X4)&(~(X10=X5)&(~(X10=X6)&(~(X10=X7)&~(X10=X8))))))))))), file('data/MPTP/proofs/3223abcad9d146c9a2641cc4d541619c.E_input', d6_enumset1)).
fof(c_0_19, plain, ![X11, X12]:k3_tarski(k2_tarski(X11,X12))=k2_xboole_0(X11,X12), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_20, plain, ![X13, X14]:k1_enumset1(X13,X13,X14)=k2_tarski(X13,X14), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_21, plain, ![X15, X16, X17, X18, X19, X20, X21, X22]:(((~r2_hidden(X18,X17)|(r2_hidden(X18,X15)|r2_hidden(X18,X16))|X17!=k2_xboole_0(X15,X16))&((~r2_hidden(X19,X15)|r2_hidden(X19,X17)|X17!=k2_xboole_0(X15,X16))&(~r2_hidden(X19,X16)|r2_hidden(X19,X17)|X17!=k2_xboole_0(X15,X16))))&(((~r2_hidden(esk1_3(X20,X21,X22),X20)|~r2_hidden(esk1_3(X20,X21,X22),X22)|X22=k2_xboole_0(X20,X21))&(~r2_hidden(esk1_3(X20,X21,X22),X21)|~r2_hidden(esk1_3(X20,X21,X22),X22)|X22=k2_xboole_0(X20,X21)))&(r2_hidden(esk1_3(X20,X21,X22),X22)|(r2_hidden(esk1_3(X20,X21,X22),X20)|r2_hidden(esk1_3(X20,X21,X22),X21))|X22=k2_xboole_0(X20,X21)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])).
cnf(c_0_22, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_24, plain, ![X24, X25, X26]:k2_enumset1(X24,X24,X25,X26)=k1_enumset1(X24,X25,X26), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_25, plain, ![X27, X28, X29, X30]:k3_enumset1(X27,X27,X28,X29,X30)=k2_enumset1(X27,X28,X29,X30), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_26, plain, ![X31, X32, X33, X34, X35]:k4_enumset1(X31,X31,X32,X33,X34,X35)=k3_enumset1(X31,X32,X33,X34,X35), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_27, plain, ![X36, X37, X38, X39, X40, X41]:k5_enumset1(X36,X36,X37,X38,X39,X40,X41)=k4_enumset1(X36,X37,X38,X39,X40,X41), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_28, plain, ![X42, X43, X44, X45, X46, X47, X48]:k6_enumset1(X42,X42,X43,X44,X45,X46,X47,X48)=k5_enumset1(X42,X43,X44,X45,X46,X47,X48), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_29, negated_conjecture, ~(![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,k1_ordinal1(X2))<=>r1_ordinal1(X1,X2))))), inference(assume_negation,[status(cth)],[t34_ordinal1])).
fof(c_0_30, plain, ![X51]:k1_ordinal1(X51)=k2_xboole_0(X51,k1_tarski(X51)), inference(variable_rename,[status(thm)],[d1_ordinal1])).
fof(c_0_31, plain, ![X52]:k2_tarski(X52,X52)=k1_tarski(X52), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_32, plain, ![X53, X54, X55, X56, X57, X58]:(((~r2_hidden(X55,X54)|X55=X53|X54!=k1_tarski(X53))&(X56!=X53|r2_hidden(X56,X54)|X54!=k1_tarski(X53)))&((~r2_hidden(esk4_2(X57,X58),X58)|esk4_2(X57,X58)!=X57|X58=k1_tarski(X57))&(r2_hidden(esk4_2(X57,X58),X58)|esk4_2(X57,X58)=X57|X58=k1_tarski(X57)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_33, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X2!=k2_xboole_0(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_34, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_35, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_36, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_37, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_38, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_39, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_40, negated_conjecture, (v3_ordinal1(esk2_0)&(v3_ordinal1(esk3_0)&((~r2_hidden(esk2_0,k1_ordinal1(esk3_0))|~r1_ordinal1(esk2_0,esk3_0))&(r2_hidden(esk2_0,k1_ordinal1(esk3_0))|r1_ordinal1(esk2_0,esk3_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])])).
cnf(c_0_41, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_42, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_43, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_44, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_45, plain, (r2_hidden(X1,X4)|r2_hidden(X1,X3)|X2!=k3_tarski(k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X4))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk2_0,k1_ordinal1(esk3_0))|r1_ordinal1(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_47, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k2_tarski(X1,X1))), inference(rw,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_48, negated_conjecture, (~r2_hidden(esk2_0,k1_ordinal1(esk3_0))|~r1_ordinal1(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_49, plain, (r2_hidden(X1,X3)|X3!=k3_tarski(k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X4))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39])).
cnf(c_0_50, plain, (X1=X3|X2!=k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_42]), c_0_23]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39])).
cnf(c_0_51, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k3_tarski(k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X2)))), inference(er,[status(thm)],[c_0_45])).
cnf(c_0_52, negated_conjecture, (r1_ordinal1(esk2_0,esk3_0)|r2_hidden(esk2_0,k3_tarski(k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_47]), c_0_23]), c_0_34]), c_0_35]), c_0_35]), c_0_36]), c_0_36]), c_0_37]), c_0_37]), c_0_38]), c_0_38]), c_0_39]), c_0_39])).
cnf(c_0_53, negated_conjecture, (~r1_ordinal1(esk2_0,esk3_0)|~r2_hidden(esk2_0,k3_tarski(k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_47]), c_0_23]), c_0_34]), c_0_35]), c_0_35]), c_0_36]), c_0_36]), c_0_37]), c_0_37]), c_0_38]), c_0_38]), c_0_39]), c_0_39])).
cnf(c_0_54, plain, (r2_hidden(X1,k3_tarski(k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)))|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_49])).
fof(c_0_55, plain, ![X60, X61]:((~r1_ordinal1(X60,X61)|r1_tarski(X60,X61)|(~v3_ordinal1(X60)|~v3_ordinal1(X61)))&(~r1_tarski(X60,X61)|r1_ordinal1(X60,X61)|(~v3_ordinal1(X60)|~v3_ordinal1(X61)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r1_ordinal1])])])).
fof(c_0_56, plain, ![X62, X63, X64]:((~v1_ordinal1(X62)|(~r2_hidden(X63,X62)|r1_tarski(X63,X62)))&((r2_hidden(esk5_1(X64),X64)|v1_ordinal1(X64))&(~r1_tarski(esk5_1(X64),X64)|v1_ordinal1(X64)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_ordinal1])])])])])])).
fof(c_0_57, plain, ![X66]:((v1_ordinal1(X66)|~v3_ordinal1(X66))&(v2_ordinal1(X66)|~v3_ordinal1(X66))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_ordinal1])])])).
cnf(c_0_58, plain, (X1=X2|~r2_hidden(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2))), inference(er,[status(thm)],[c_0_50])).
cnf(c_0_59, negated_conjecture, (r1_ordinal1(esk2_0,esk3_0)|r2_hidden(esk2_0,k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))|r2_hidden(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_60, negated_conjecture, (~r1_ordinal1(esk2_0,esk3_0)|~r2_hidden(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_61, plain, (r1_ordinal1(X1,X2)|~r1_tarski(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_62, negated_conjecture, (v3_ordinal1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_63, negated_conjecture, (v3_ordinal1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_64, plain, (r1_tarski(X2,X1)|~v1_ordinal1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_65, plain, (v1_ordinal1(X1)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_66, plain, (r1_tarski(X1,X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_67, negated_conjecture, (esk2_0=esk3_0|r1_ordinal1(esk2_0,esk3_0)|r2_hidden(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_68, negated_conjecture, (~r1_tarski(esk2_0,esk3_0)|~r2_hidden(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62]), c_0_63])])).
cnf(c_0_69, plain, (r1_tarski(X1,X2)|~v3_ordinal1(X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_70, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_71, plain, ![X67, X68]:(((r1_tarski(X67,X68)|X67!=X68)&(r1_tarski(X68,X67)|X67!=X68))&(~r1_tarski(X67,X68)|~r1_tarski(X68,X67)|X67=X68)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_72, negated_conjecture, (esk2_0=esk3_0|r1_tarski(esk2_0,esk3_0)|r2_hidden(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_62]), c_0_63])])).
cnf(c_0_73, negated_conjecture, (~r2_hidden(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_62])])).
fof(c_0_74, plain, ![X69, X70]:(~v3_ordinal1(X69)|(~v3_ordinal1(X70)|(r1_ordinal1(X69,X70)|r2_hidden(X70,X69)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t26_ordinal1])])])).
fof(c_0_75, plain, ![X71, X72]:(~v3_ordinal1(X71)|~v3_ordinal1(X72)|(r1_ordinal1(X71,X72)|r1_ordinal1(X72,X71))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[connectedness_r1_ordinal1])])).
cnf(c_0_76, plain, (r2_hidden(X1,X3)|X3!=k3_tarski(k6_enumset1(X4,X4,X4,X4,X4,X4,X4,X2))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39])).
cnf(c_0_77, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_78, negated_conjecture, (esk2_0=esk3_0|r1_tarski(esk2_0,esk3_0)), inference(sr,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_79, plain, (r1_ordinal1(X1,X2)|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_80, plain, (r1_ordinal1(X1,X2)|r1_ordinal1(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_75])).
fof(c_0_81, plain, ![X73, X74, X75, X76, X77, X78, X79, X80, X81, X82, X83, X84, X85, X86, X87, X88, X89, X90, X91, X92]:(((~r2_hidden(X82,X81)|(X82=X73|(X82=X74|(X82=X75|(X82=X76|(X82=X77|(X82=X78|(X82=X79|X82=X80)))))))|X81!=k6_enumset1(X73,X74,X75,X76,X77,X78,X79,X80))&((X83!=X73|r2_hidden(X83,X81)|X81!=k6_enumset1(X73,X74,X75,X76,X77,X78,X79,X80))&((X83!=X74|r2_hidden(X83,X81)|X81!=k6_enumset1(X73,X74,X75,X76,X77,X78,X79,X80))&((X83!=X75|r2_hidden(X83,X81)|X81!=k6_enumset1(X73,X74,X75,X76,X77,X78,X79,X80))&((X83!=X76|r2_hidden(X83,X81)|X81!=k6_enumset1(X73,X74,X75,X76,X77,X78,X79,X80))&((X83!=X77|r2_hidden(X83,X81)|X81!=k6_enumset1(X73,X74,X75,X76,X77,X78,X79,X80))&((X83!=X78|r2_hidden(X83,X81)|X81!=k6_enumset1(X73,X74,X75,X76,X77,X78,X79,X80))&((X83!=X79|r2_hidden(X83,X81)|X81!=k6_enumset1(X73,X74,X75,X76,X77,X78,X79,X80))&(X83!=X80|r2_hidden(X83,X81)|X81!=k6_enumset1(X73,X74,X75,X76,X77,X78,X79,X80))))))))))&(((esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92)!=X84|~r2_hidden(esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92),X92)|X92=k6_enumset1(X84,X85,X86,X87,X88,X89,X90,X91))&((esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92)!=X85|~r2_hidden(esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92),X92)|X92=k6_enumset1(X84,X85,X86,X87,X88,X89,X90,X91))&((esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92)!=X86|~r2_hidden(esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92),X92)|X92=k6_enumset1(X84,X85,X86,X87,X88,X89,X90,X91))&((esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92)!=X87|~r2_hidden(esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92),X92)|X92=k6_enumset1(X84,X85,X86,X87,X88,X89,X90,X91))&((esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92)!=X88|~r2_hidden(esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92),X92)|X92=k6_enumset1(X84,X85,X86,X87,X88,X89,X90,X91))&((esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92)!=X89|~r2_hidden(esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92),X92)|X92=k6_enumset1(X84,X85,X86,X87,X88,X89,X90,X91))&((esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92)!=X90|~r2_hidden(esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92),X92)|X92=k6_enumset1(X84,X85,X86,X87,X88,X89,X90,X91))&(esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92)!=X91|~r2_hidden(esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92),X92)|X92=k6_enumset1(X84,X85,X86,X87,X88,X89,X90,X91)))))))))&(r2_hidden(esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92),X92)|(esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92)=X84|(esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92)=X85|(esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92)=X86|(esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92)=X87|(esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92)=X88|(esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92)=X89|(esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92)=X90|esk6_9(X84,X85,X86,X87,X88,X89,X90,X91,X92)=X91)))))))|X92=k6_enumset1(X84,X85,X86,X87,X88,X89,X90,X91)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_enumset1])])])])])])])).
cnf(c_0_82, plain, (r2_hidden(X1,k3_tarski(k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_76])).
cnf(c_0_83, negated_conjecture, (esk2_0=esk3_0|~r1_tarski(esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_84, plain, (r1_tarski(X1,X2)|r2_hidden(X2,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_66, c_0_79])).
cnf(c_0_85, negated_conjecture, (r1_ordinal1(X1,esk3_0)|r1_ordinal1(esk3_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_80, c_0_62])).
cnf(c_0_86, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k6_enumset1(X2,X4,X5,X6,X7,X8,X9,X10)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_87, negated_conjecture, (~r1_ordinal1(esk2_0,esk3_0)|~r2_hidden(esk2_0,k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))), inference(spm,[status(thm)],[c_0_53, c_0_82])).
cnf(c_0_88, negated_conjecture, (esk2_0=esk3_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_63]), c_0_62])]), c_0_73])).
cnf(c_0_89, negated_conjecture, (r1_ordinal1(esk3_0,esk3_0)), inference(spm,[status(thm)],[c_0_85, c_0_62])).
cnf(c_0_90, plain, (r2_hidden(X1,k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_86])])).
cnf(c_0_91, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_88]), c_0_89]), c_0_88]), c_0_90])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 92
# Proof object clause steps            : 53
# Proof object formula steps           : 39
# Proof object conjectures             : 22
# Proof object clause conjectures      : 19
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 77
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 53
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 44
# Processed clauses                    : 140
# ...of these trivial                  : 0
# ...subsumed                          : 11
# ...remaining for further processing  : 129
# Other redundant clauses eliminated   : 43
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 24
# Generated clauses                    : 261
# ...of the previous two non-trivial   : 229
# Contextual simplify-reflections      : 0
# Paramodulations                      : 219
# Factorizations                       : 5
# Equation resolutions                 : 43
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 43
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 31
# Current number of unprocessed clauses: 167
# ...number of literals in the above   : 751
# Current number of archived formulas  : 0
# Current number of archived clauses   : 79
# Clause-clause subsumption calls (NU) : 512
# Rec. Clause-clause subsumption calls : 328
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 46
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 77
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7226

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.000 s
# Total time               : 0.015 s
# Maximum resident set size: 28748 pages
