# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/326c3d2ed6b14c8d882eea1ee3d998a6.E_input', d2_subset_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/326c3d2ed6b14c8d882eea1ee3d998a6.E_input', d1_xboole_0)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/326c3d2ed6b14c8d882eea1ee3d998a6.E_input', d9_lattice3)).
fof(d5_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>(v2_yellow_0(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&r2_lattice3(X1,u1_struct_0(X1),X2)))), file('data/MPTP/proofs/326c3d2ed6b14c8d882eea1ee3d998a6.E_input', d5_yellow_0)).
fof(t16_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3))))))), file('data/MPTP/proofs/326c3d2ed6b14c8d882eea1ee3d998a6.E_input', t16_yellow_0)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/326c3d2ed6b14c8d882eea1ee3d998a6.E_input', t2_subset)).
fof(t6_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_lattice3(X1,k1_xboole_0,X2)&r1_lattice3(X1,k1_xboole_0,X2)))), file('data/MPTP/proofs/326c3d2ed6b14c8d882eea1ee3d998a6.E_input', t6_yellow_0)).
fof(t43_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v2_yellow_0(X1)&l1_orders_2(X1))))=>(r2_yellow_0(X1,k1_xboole_0)&r1_yellow_0(X1,u1_struct_0(X1)))), file('data/MPTP/proofs/326c3d2ed6b14c8d882eea1ee3d998a6.E_input', t43_yellow_0)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/326c3d2ed6b14c8d882eea1ee3d998a6.E_input', d8_lattice3)).
fof(d8_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X5)=>r1_orders_2(X1,X5,X4))))=>X4=X3))))))), file('data/MPTP/proofs/326c3d2ed6b14c8d882eea1ee3d998a6.E_input', d8_yellow_0)).
fof(d4_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>(v1_yellow_0(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&r1_lattice3(X1,u1_struct_0(X1),X2)))), file('data/MPTP/proofs/326c3d2ed6b14c8d882eea1ee3d998a6.E_input', d4_yellow_0)).
fof(t30_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4)))))&((r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))=>(X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3)))))), file('data/MPTP/proofs/326c3d2ed6b14c8d882eea1ee3d998a6.E_input', t30_yellow_0)).
fof(t42_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))=>(r1_yellow_0(X1,k1_xboole_0)&r2_yellow_0(X1,u1_struct_0(X1)))), file('data/MPTP/proofs/326c3d2ed6b14c8d882eea1ee3d998a6.E_input', t42_yellow_0)).
fof(t15_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4))))))), file('data/MPTP/proofs/326c3d2ed6b14c8d882eea1ee3d998a6.E_input', t15_yellow_0)).
fof(c_0_14, plain, ![X60, X61]:(((~m1_subset_1(X61,X60)|r2_hidden(X61,X60)|v1_xboole_0(X60))&(~r2_hidden(X61,X60)|m1_subset_1(X61,X60)|v1_xboole_0(X60)))&((~m1_subset_1(X61,X60)|v1_xboole_0(X61)|~v1_xboole_0(X60))&(~v1_xboole_0(X61)|m1_subset_1(X61,X60)|~v1_xboole_0(X60)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_15, plain, ![X36, X37, X38]:((~v1_xboole_0(X36)|~r2_hidden(X37,X36))&(r2_hidden(esk8_1(X38),X38)|v1_xboole_0(X38))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_16, plain, ![X20, X21, X22, X23]:((~r2_lattice3(X20,X21,X22)|(~m1_subset_1(X23,u1_struct_0(X20))|(~r2_hidden(X23,X21)|r1_orders_2(X20,X23,X22)))|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))&((m1_subset_1(esk5_3(X20,X21,X22),u1_struct_0(X20))|r2_lattice3(X20,X21,X22)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))&((r2_hidden(esk5_3(X20,X21,X22),X21)|r2_lattice3(X20,X21,X22)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))&(~r1_orders_2(X20,esk5_3(X20,X21,X22),X22)|r2_lattice3(X20,X21,X22)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
fof(c_0_17, plain, ![X8, X10]:(((m1_subset_1(esk2_1(X8),u1_struct_0(X8))|~v2_yellow_0(X8)|~l1_orders_2(X8))&(r2_lattice3(X8,u1_struct_0(X8),esk2_1(X8))|~v2_yellow_0(X8)|~l1_orders_2(X8)))&(~m1_subset_1(X10,u1_struct_0(X8))|~r2_lattice3(X8,u1_struct_0(X8),X10)|v2_yellow_0(X8)|~l1_orders_2(X8))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_yellow_0])])])])])).
cnf(c_0_18, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_20, plain, ![X11, X12, X14, X15, X16]:(((m1_subset_1(esk3_2(X11,X12),u1_struct_0(X11))|~r2_yellow_0(X11,X12)|(~v5_orders_2(X11)|~l1_orders_2(X11)))&((r1_lattice3(X11,X12,esk3_2(X11,X12))|~r2_yellow_0(X11,X12)|(~v5_orders_2(X11)|~l1_orders_2(X11)))&(~m1_subset_1(X14,u1_struct_0(X11))|(~r1_lattice3(X11,X12,X14)|r1_orders_2(X11,X14,esk3_2(X11,X12)))|~r2_yellow_0(X11,X12)|(~v5_orders_2(X11)|~l1_orders_2(X11)))))&((m1_subset_1(esk4_3(X11,X15,X16),u1_struct_0(X11))|~r1_lattice3(X11,X15,X16)|~m1_subset_1(X16,u1_struct_0(X11))|r2_yellow_0(X11,X15)|(~v5_orders_2(X11)|~l1_orders_2(X11)))&((r1_lattice3(X11,X15,esk4_3(X11,X15,X16))|~r1_lattice3(X11,X15,X16)|~m1_subset_1(X16,u1_struct_0(X11))|r2_yellow_0(X11,X15)|(~v5_orders_2(X11)|~l1_orders_2(X11)))&(~r1_orders_2(X11,esk4_3(X11,X15,X16),X16)|~r1_lattice3(X11,X15,X16)|~m1_subset_1(X16,u1_struct_0(X11))|r2_yellow_0(X11,X15)|(~v5_orders_2(X11)|~l1_orders_2(X11)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_yellow_0])])])])])])).
cnf(c_0_21, plain, (r1_orders_2(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (r2_lattice3(X1,u1_struct_0(X1),esk2_1(X1))|~v2_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_24, plain, (m1_subset_1(esk2_1(X1),u1_struct_0(X1))|~v2_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (r2_yellow_0(X1,X2)|~r1_orders_2(X1,esk4_3(X1,X2,X3),X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (r1_orders_2(X1,X2,esk2_1(X1))|~r2_hidden(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_yellow_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24])).
fof(c_0_27, plain, ![X51, X52]:(~m1_subset_1(X51,X52)|(v1_xboole_0(X52)|r2_hidden(X51,X52))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_28, plain, (r2_yellow_0(X1,X2)|~r2_hidden(esk4_3(X1,X2,esk2_1(X1)),u1_struct_0(X1))|~r1_lattice3(X1,X2,esk2_1(X1))|~l1_orders_2(X1)|~v2_yellow_0(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_24])).
cnf(c_0_29, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_30, plain, (v1_xboole_0(u1_struct_0(X1))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,esk2_1(X1))|~m1_subset_1(esk4_3(X1,X2,esk2_1(X1)),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_yellow_0(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_31, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X1))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_32, plain, ![X18, X19]:((r2_lattice3(X18,k1_xboole_0,X19)|~m1_subset_1(X19,u1_struct_0(X18))|~l1_orders_2(X18))&(r1_lattice3(X18,k1_xboole_0,X19)|~m1_subset_1(X19,u1_struct_0(X18))|~l1_orders_2(X18))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_yellow_0])])])])).
fof(c_0_33, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v2_yellow_0(X1)&l1_orders_2(X1))))=>(r2_yellow_0(X1,k1_xboole_0)&r1_yellow_0(X1,u1_struct_0(X1))))), inference(assume_negation,[status(cth)],[t43_yellow_0])).
fof(c_0_34, plain, ![X31, X32, X33, X34]:((~r1_lattice3(X31,X32,X33)|(~m1_subset_1(X34,u1_struct_0(X31))|(~r2_hidden(X34,X32)|r1_orders_2(X31,X33,X34)))|~m1_subset_1(X33,u1_struct_0(X31))|~l1_orders_2(X31))&((m1_subset_1(esk7_3(X31,X32,X33),u1_struct_0(X31))|r1_lattice3(X31,X32,X33)|~m1_subset_1(X33,u1_struct_0(X31))|~l1_orders_2(X31))&((r2_hidden(esk7_3(X31,X32,X33),X32)|r1_lattice3(X31,X32,X33)|~m1_subset_1(X33,u1_struct_0(X31))|~l1_orders_2(X31))&(~r1_orders_2(X31,X33,esk7_3(X31,X32,X33))|r1_lattice3(X31,X32,X33)|~m1_subset_1(X33,u1_struct_0(X31))|~l1_orders_2(X31))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
cnf(c_0_35, plain, (v1_xboole_0(u1_struct_0(X1))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,esk2_1(X1))|~l1_orders_2(X1)|~v2_yellow_0(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_24])).
cnf(c_0_36, plain, (r1_lattice3(X1,k1_xboole_0,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_37, negated_conjecture, ((~v2_struct_0(esk1_0)&(v5_orders_2(esk1_0)&(v2_yellow_0(esk1_0)&l1_orders_2(esk1_0))))&(~r2_yellow_0(esk1_0,k1_xboole_0)|~r1_yellow_0(esk1_0,u1_struct_0(esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_33])])])])).
fof(c_0_38, plain, ![X40, X41, X43, X44, X46, X47, X50]:(((m1_subset_1(esk9_2(X40,X41),u1_struct_0(X40))|~r2_yellow_0(X40,X41)|~l1_orders_2(X40))&((r1_lattice3(X40,X41,esk9_2(X40,X41))|~r2_yellow_0(X40,X41)|~l1_orders_2(X40))&((~m1_subset_1(X43,u1_struct_0(X40))|(~r1_lattice3(X40,X41,X43)|r1_orders_2(X40,X43,esk9_2(X40,X41)))|~r2_yellow_0(X40,X41)|~l1_orders_2(X40))&((m1_subset_1(esk10_3(X40,X41,X44),u1_struct_0(X40))|~r1_lattice3(X40,X41,X44)|X44=esk9_2(X40,X41)|~m1_subset_1(X44,u1_struct_0(X40))|~r2_yellow_0(X40,X41)|~l1_orders_2(X40))&((r1_lattice3(X40,X41,esk10_3(X40,X41,X44))|~r1_lattice3(X40,X41,X44)|X44=esk9_2(X40,X41)|~m1_subset_1(X44,u1_struct_0(X40))|~r2_yellow_0(X40,X41)|~l1_orders_2(X40))&(~r1_orders_2(X40,esk10_3(X40,X41,X44),X44)|~r1_lattice3(X40,X41,X44)|X44=esk9_2(X40,X41)|~m1_subset_1(X44,u1_struct_0(X40))|~r2_yellow_0(X40,X41)|~l1_orders_2(X40)))))))&(((m1_subset_1(esk12_3(X40,X46,X47),u1_struct_0(X40))|m1_subset_1(esk11_3(X40,X46,X47),u1_struct_0(X40))|~r1_lattice3(X40,X46,X47)|~m1_subset_1(X47,u1_struct_0(X40))|r2_yellow_0(X40,X46)|~l1_orders_2(X40))&(((r1_lattice3(X40,X46,esk12_3(X40,X46,X47))|m1_subset_1(esk11_3(X40,X46,X47),u1_struct_0(X40))|~r1_lattice3(X40,X46,X47)|~m1_subset_1(X47,u1_struct_0(X40))|r2_yellow_0(X40,X46)|~l1_orders_2(X40))&(~m1_subset_1(X50,u1_struct_0(X40))|(~r1_lattice3(X40,X46,X50)|r1_orders_2(X40,X50,esk12_3(X40,X46,X47)))|m1_subset_1(esk11_3(X40,X46,X47),u1_struct_0(X40))|~r1_lattice3(X40,X46,X47)|~m1_subset_1(X47,u1_struct_0(X40))|r2_yellow_0(X40,X46)|~l1_orders_2(X40)))&(esk12_3(X40,X46,X47)!=X47|m1_subset_1(esk11_3(X40,X46,X47),u1_struct_0(X40))|~r1_lattice3(X40,X46,X47)|~m1_subset_1(X47,u1_struct_0(X40))|r2_yellow_0(X40,X46)|~l1_orders_2(X40))))&(((m1_subset_1(esk12_3(X40,X46,X47),u1_struct_0(X40))|r1_lattice3(X40,X46,esk11_3(X40,X46,X47))|~r1_lattice3(X40,X46,X47)|~m1_subset_1(X47,u1_struct_0(X40))|r2_yellow_0(X40,X46)|~l1_orders_2(X40))&(((r1_lattice3(X40,X46,esk12_3(X40,X46,X47))|r1_lattice3(X40,X46,esk11_3(X40,X46,X47))|~r1_lattice3(X40,X46,X47)|~m1_subset_1(X47,u1_struct_0(X40))|r2_yellow_0(X40,X46)|~l1_orders_2(X40))&(~m1_subset_1(X50,u1_struct_0(X40))|(~r1_lattice3(X40,X46,X50)|r1_orders_2(X40,X50,esk12_3(X40,X46,X47)))|r1_lattice3(X40,X46,esk11_3(X40,X46,X47))|~r1_lattice3(X40,X46,X47)|~m1_subset_1(X47,u1_struct_0(X40))|r2_yellow_0(X40,X46)|~l1_orders_2(X40)))&(esk12_3(X40,X46,X47)!=X47|r1_lattice3(X40,X46,esk11_3(X40,X46,X47))|~r1_lattice3(X40,X46,X47)|~m1_subset_1(X47,u1_struct_0(X40))|r2_yellow_0(X40,X46)|~l1_orders_2(X40))))&((m1_subset_1(esk12_3(X40,X46,X47),u1_struct_0(X40))|~r1_orders_2(X40,esk11_3(X40,X46,X47),X47)|~r1_lattice3(X40,X46,X47)|~m1_subset_1(X47,u1_struct_0(X40))|r2_yellow_0(X40,X46)|~l1_orders_2(X40))&(((r1_lattice3(X40,X46,esk12_3(X40,X46,X47))|~r1_orders_2(X40,esk11_3(X40,X46,X47),X47)|~r1_lattice3(X40,X46,X47)|~m1_subset_1(X47,u1_struct_0(X40))|r2_yellow_0(X40,X46)|~l1_orders_2(X40))&(~m1_subset_1(X50,u1_struct_0(X40))|(~r1_lattice3(X40,X46,X50)|r1_orders_2(X40,X50,esk12_3(X40,X46,X47)))|~r1_orders_2(X40,esk11_3(X40,X46,X47),X47)|~r1_lattice3(X40,X46,X47)|~m1_subset_1(X47,u1_struct_0(X40))|r2_yellow_0(X40,X46)|~l1_orders_2(X40)))&(esk12_3(X40,X46,X47)!=X47|~r1_orders_2(X40,esk11_3(X40,X46,X47),X47)|~r1_lattice3(X40,X46,X47)|~m1_subset_1(X47,u1_struct_0(X40))|r2_yellow_0(X40,X46)|~l1_orders_2(X40))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_yellow_0])])])])])])).
cnf(c_0_39, plain, (r2_hidden(esk7_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_40, plain, (v1_xboole_0(u1_struct_0(X1))|r2_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)|~v2_yellow_0(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_24])).
cnf(c_0_41, negated_conjecture, (v2_yellow_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_42, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_43, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_44, plain, (r1_orders_2(X2,X1,esk9_2(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~r2_yellow_0(X2,X3)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_45, plain, (m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_46, plain, (r1_lattice3(X1,X2,X3)|~v1_xboole_0(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_19, c_0_39])).
cnf(c_0_47, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_yellow_0(esk1_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42]), c_0_43])])).
fof(c_0_48, plain, ![X62, X64]:(((m1_subset_1(esk15_1(X62),u1_struct_0(X62))|~v1_yellow_0(X62)|~l1_orders_2(X62))&(r1_lattice3(X62,u1_struct_0(X62),esk15_1(X62))|~v1_yellow_0(X62)|~l1_orders_2(X62)))&(~m1_subset_1(X64,u1_struct_0(X62))|~r1_lattice3(X62,u1_struct_0(X62),X64)|v1_yellow_0(X62)|~l1_orders_2(X62))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_yellow_0])])])])])).
fof(c_0_49, plain, ![X25, X26, X27, X28, X29]:(((r2_lattice3(X25,X27,X26)|(X26!=k1_yellow_0(X25,X27)|~r1_yellow_0(X25,X27))|~m1_subset_1(X26,u1_struct_0(X25))|(~v5_orders_2(X25)|~l1_orders_2(X25)))&(~m1_subset_1(X28,u1_struct_0(X25))|(~r2_lattice3(X25,X27,X28)|r1_orders_2(X25,X26,X28))|(X26!=k1_yellow_0(X25,X27)|~r1_yellow_0(X25,X27))|~m1_subset_1(X26,u1_struct_0(X25))|(~v5_orders_2(X25)|~l1_orders_2(X25))))&(((X26=k1_yellow_0(X25,X29)|(m1_subset_1(esk6_3(X25,X26,X29),u1_struct_0(X25))|~r2_lattice3(X25,X29,X26))|~m1_subset_1(X26,u1_struct_0(X25))|(~v5_orders_2(X25)|~l1_orders_2(X25)))&(r1_yellow_0(X25,X29)|(m1_subset_1(esk6_3(X25,X26,X29),u1_struct_0(X25))|~r2_lattice3(X25,X29,X26))|~m1_subset_1(X26,u1_struct_0(X25))|(~v5_orders_2(X25)|~l1_orders_2(X25))))&(((X26=k1_yellow_0(X25,X29)|(r2_lattice3(X25,X29,esk6_3(X25,X26,X29))|~r2_lattice3(X25,X29,X26))|~m1_subset_1(X26,u1_struct_0(X25))|(~v5_orders_2(X25)|~l1_orders_2(X25)))&(r1_yellow_0(X25,X29)|(r2_lattice3(X25,X29,esk6_3(X25,X26,X29))|~r2_lattice3(X25,X29,X26))|~m1_subset_1(X26,u1_struct_0(X25))|(~v5_orders_2(X25)|~l1_orders_2(X25))))&((X26=k1_yellow_0(X25,X29)|(~r1_orders_2(X25,X26,esk6_3(X25,X26,X29))|~r2_lattice3(X25,X29,X26))|~m1_subset_1(X26,u1_struct_0(X25))|(~v5_orders_2(X25)|~l1_orders_2(X25)))&(r1_yellow_0(X25,X29)|(~r1_orders_2(X25,X26,esk6_3(X25,X26,X29))|~r2_lattice3(X25,X29,X26))|~m1_subset_1(X26,u1_struct_0(X25))|(~v5_orders_2(X25)|~l1_orders_2(X25))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_yellow_0])])])])])])).
cnf(c_0_50, plain, (r2_yellow_0(X1,X2)|~r1_lattice3(X1,X3,esk4_3(X1,X2,esk9_2(X1,X3)))|~r1_lattice3(X1,X2,esk9_2(X1,X3))|~m1_subset_1(esk4_3(X1,X2,esk9_2(X1,X3)),u1_struct_0(X1))|~r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_44]), c_0_45])).
cnf(c_0_51, negated_conjecture, (r1_lattice3(X1,u1_struct_0(esk1_0),X2)|r2_yellow_0(esk1_0,k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_52, plain, (v1_yellow_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,u1_struct_0(X2),X1)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_53, plain, (r1_yellow_0(X1,X2)|r2_lattice3(X1,X2,esk6_3(X1,X3,X2))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_54, plain, (r1_yellow_0(X1,X2)|m1_subset_1(esk6_3(X1,X3,X2),u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_55, negated_conjecture, (r2_yellow_0(esk1_0,k1_xboole_0)|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,esk9_2(X1,u1_struct_0(esk1_0)))|~m1_subset_1(esk4_3(X1,X2,esk9_2(X1,u1_struct_0(esk1_0))),u1_struct_0(X1))|~r2_yellow_0(X1,u1_struct_0(esk1_0))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
fof(c_0_56, plain, ![X7]:((r1_yellow_0(X7,k1_xboole_0)|(v2_struct_0(X7)|(~v5_orders_2(X7)|(~v1_yellow_0(X7)|~l1_orders_2(X7)))))&(r2_yellow_0(X7,u1_struct_0(X7))|(v2_struct_0(X7)|(~v5_orders_2(X7)|(~v1_yellow_0(X7)|~l1_orders_2(X7)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t42_yellow_0])])])])).
cnf(c_0_57, negated_conjecture, (v1_yellow_0(esk1_0)|r2_yellow_0(esk1_0,k1_xboole_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_51]), c_0_42])])).
cnf(c_0_58, plain, (r1_yellow_0(X1,X2)|~r1_orders_2(X1,X3,esk6_3(X1,X3,X2))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_59, plain, (r1_orders_2(X1,X2,esk6_3(X1,X3,X4))|r1_yellow_0(X1,X4)|~r2_hidden(X2,X4)|~r2_lattice3(X1,X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_53]), c_0_54])).
cnf(c_0_60, negated_conjecture, (r2_yellow_0(esk1_0,k1_xboole_0)|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,esk9_2(X1,u1_struct_0(esk1_0)))|~r2_yellow_0(X1,u1_struct_0(esk1_0))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_31]), c_0_45])).
cnf(c_0_61, plain, (r2_yellow_0(X1,u1_struct_0(X1))|v2_struct_0(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_62, negated_conjecture, (v1_yellow_0(esk1_0)|r2_yellow_0(esk1_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_24]), c_0_42]), c_0_41])])).
cnf(c_0_63, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_64, plain, (r1_yellow_0(X1,X2)|~r2_hidden(X3,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_65, negated_conjecture, (r2_yellow_0(esk1_0,k1_xboole_0)|r2_yellow_0(X1,k1_xboole_0)|~r2_yellow_0(X1,u1_struct_0(esk1_0))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_36]), c_0_45])).
cnf(c_0_66, negated_conjecture, (r2_yellow_0(esk1_0,u1_struct_0(esk1_0))|r2_yellow_0(esk1_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_42]), c_0_43])]), c_0_63])).
cnf(c_0_67, plain, (r1_yellow_0(X1,u1_struct_0(X1))|~r2_hidden(esk2_1(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_yellow_0(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_22]), c_0_23])).
cnf(c_0_68, negated_conjecture, (~r2_yellow_0(esk1_0,k1_xboole_0)|~r1_yellow_0(esk1_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_69, negated_conjecture, (r2_yellow_0(esk1_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_42]), c_0_43])])).
fof(c_0_70, plain, ![X53, X54, X56, X57, X58]:(((m1_subset_1(esk13_2(X53,X54),u1_struct_0(X53))|~r1_yellow_0(X53,X54)|(~v5_orders_2(X53)|~l1_orders_2(X53)))&((r2_lattice3(X53,X54,esk13_2(X53,X54))|~r1_yellow_0(X53,X54)|(~v5_orders_2(X53)|~l1_orders_2(X53)))&(~m1_subset_1(X56,u1_struct_0(X53))|(~r2_lattice3(X53,X54,X56)|r1_orders_2(X53,esk13_2(X53,X54),X56))|~r1_yellow_0(X53,X54)|(~v5_orders_2(X53)|~l1_orders_2(X53)))))&((m1_subset_1(esk14_3(X53,X57,X58),u1_struct_0(X53))|~r2_lattice3(X53,X57,X58)|~m1_subset_1(X58,u1_struct_0(X53))|r1_yellow_0(X53,X57)|(~v5_orders_2(X53)|~l1_orders_2(X53)))&((r2_lattice3(X53,X57,esk14_3(X53,X57,X58))|~r2_lattice3(X53,X57,X58)|~m1_subset_1(X58,u1_struct_0(X53))|r1_yellow_0(X53,X57)|(~v5_orders_2(X53)|~l1_orders_2(X53)))&(~r1_orders_2(X53,X58,esk14_3(X53,X57,X58))|~r2_lattice3(X53,X57,X58)|~m1_subset_1(X58,u1_struct_0(X53))|r1_yellow_0(X53,X57)|(~v5_orders_2(X53)|~l1_orders_2(X53)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_yellow_0])])])])])])).
cnf(c_0_71, plain, (v1_xboole_0(u1_struct_0(X1))|r1_yellow_0(X1,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_yellow_0(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_29]), c_0_24])).
cnf(c_0_72, negated_conjecture, (~r1_yellow_0(esk1_0,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_69])])).
cnf(c_0_73, plain, (r1_orders_2(X2,esk13_2(X2,X3),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r1_yellow_0(X2,X3)|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_74, plain, (m1_subset_1(esk13_2(X1,X2),u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_75, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_41]), c_0_42]), c_0_43])]), c_0_72])).
cnf(c_0_76, plain, (r1_yellow_0(X1,X2)|~r2_lattice3(X1,X3,esk6_3(X1,esk13_2(X1,X3),X2))|~r2_lattice3(X1,X2,esk13_2(X1,X3))|~m1_subset_1(esk6_3(X1,esk13_2(X1,X3),X2),u1_struct_0(X1))|~r1_yellow_0(X1,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_73]), c_0_74])).
cnf(c_0_77, plain, (r2_lattice3(X1,k1_xboole_0,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_78, plain, (r2_hidden(esk5_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_79, negated_conjecture, (r1_lattice3(X1,u1_struct_0(esk1_0),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_46, c_0_75])).
cnf(c_0_80, plain, (r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,esk13_2(X1,k1_xboole_0))|~m1_subset_1(esk6_3(X1,esk13_2(X1,k1_xboole_0),X2),u1_struct_0(X1))|~r1_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_81, plain, (r2_lattice3(X1,X2,X3)|~v1_xboole_0(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_19, c_0_78])).
cnf(c_0_82, plain, (r1_yellow_0(X1,k1_xboole_0)|v2_struct_0(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_83, negated_conjecture, (v1_yellow_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_79]), c_0_42])])).
cnf(c_0_84, plain, (r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,esk13_2(X1,k1_xboole_0))|~r1_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_54]), c_0_74])).
cnf(c_0_85, negated_conjecture, (r2_lattice3(X1,u1_struct_0(esk1_0),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_81, c_0_75])).
cnf(c_0_86, negated_conjecture, (r1_yellow_0(esk1_0,k1_xboole_0)|~v1_yellow_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_82]), c_0_42]), c_0_43])])).
cnf(c_0_87, negated_conjecture, (v1_yellow_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_24]), c_0_42]), c_0_41])])).
cnf(c_0_88, negated_conjecture, (r1_yellow_0(X1,u1_struct_0(esk1_0))|~r1_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_74])).
cnf(c_0_89, negated_conjecture, (r1_yellow_0(esk1_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_87])])).
cnf(c_0_90, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_88]), c_0_89]), c_0_42]), c_0_43])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 91
# Proof object clause steps            : 62
# Proof object formula steps           : 29
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 32
# Proof object simplifying inferences  : 50
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 68
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 68
# Processed clauses                    : 498
# ...of these trivial                  : 1
# ...subsumed                          : 80
# ...remaining for further processing  : 417
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 23
# Backward-rewritten                   : 37
# Generated clauses                    : 1198
# ...of the previous two non-trivial   : 1112
# Contextual simplify-reflections      : 105
# Paramodulations                      : 1196
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 288
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 277
# Current number of unprocessed clauses: 677
# ...number of literals in the above   : 6150
# Current number of archived formulas  : 0
# Current number of archived clauses   : 127
# Clause-clause subsumption calls (NU) : 34560
# Rec. Clause-clause subsumption calls : 1324
# Non-unit clause-clause subsumptions  : 207
# Unit Clause-clause subsumption calls : 656
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 40646

# -------------------------------------------------
# User time                : 0.040 s
# System time              : 0.007 s
# Total time               : 0.046 s
# Maximum resident set size: 28852 pages
