# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t78_tops_2, conjecture, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(v1_tops_2(X2,X1)<=>r1_tarski(X2,u1_pre_topc(X1))))), file('data/MPTP/proofs/32e2640039eb4ca49a48d80e2e4f158c.E_input', t78_tops_2)).
fof(t7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,X2)=>r2_hidden(X4,X3)))=>r1_tarski(X2,X3)))), file('data/MPTP/proofs/32e2640039eb4ca49a48d80e2e4f158c.E_input', t7_subset_1)).
fof(d5_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>r2_hidden(X2,u1_pre_topc(X1))))), file('data/MPTP/proofs/32e2640039eb4ca49a48d80e2e4f158c.E_input', d5_pre_topc)).
fof(d1_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(v1_tops_2(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X3,X2)=>v3_pre_topc(X3,X1)))))), file('data/MPTP/proofs/32e2640039eb4ca49a48d80e2e4f158c.E_input', d1_tops_2)).
fof(t18_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>((r1_tarski(X2,X3)&v1_tops_2(X3,X1))=>v1_tops_2(X2,X1))))), file('data/MPTP/proofs/32e2640039eb4ca49a48d80e2e4f158c.E_input', t18_tops_2)).
fof(dt_u1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/32e2640039eb4ca49a48d80e2e4f158c.E_input', dt_u1_pre_topc)).
fof(c_0_6, negated_conjecture, ~(![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(v1_tops_2(X2,X1)<=>r1_tarski(X2,u1_pre_topc(X1)))))), inference(assume_negation,[status(cth)],[t78_tops_2])).
fof(c_0_7, plain, ![X7, X8, X9]:((m1_subset_1(esk3_3(X7,X8,X9),X7)|r1_tarski(X8,X9)|~m1_subset_1(X9,k1_zfmisc_1(X7))|~m1_subset_1(X8,k1_zfmisc_1(X7)))&((r2_hidden(esk3_3(X7,X8,X9),X8)|r1_tarski(X8,X9)|~m1_subset_1(X9,k1_zfmisc_1(X7))|~m1_subset_1(X8,k1_zfmisc_1(X7)))&(~r2_hidden(esk3_3(X7,X8,X9),X9)|r1_tarski(X8,X9)|~m1_subset_1(X9,k1_zfmisc_1(X7))|~m1_subset_1(X8,k1_zfmisc_1(X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_subset_1])])])])])).
fof(c_0_8, plain, ![X11, X12]:((~v3_pre_topc(X12,X11)|r2_hidden(X12,u1_pre_topc(X11))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))&(~r2_hidden(X12,u1_pre_topc(X11))|v3_pre_topc(X12,X11)|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_pre_topc])])])])).
fof(c_0_9, plain, ![X13, X14, X15]:((~v1_tops_2(X14,X13)|(~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X13)))|(~r2_hidden(X15,X14)|v3_pre_topc(X15,X13)))|~m1_subset_1(X14,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X13))))|~l1_pre_topc(X13))&((m1_subset_1(esk4_2(X13,X14),k1_zfmisc_1(u1_struct_0(X13)))|v1_tops_2(X14,X13)|~m1_subset_1(X14,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X13))))|~l1_pre_topc(X13))&((r2_hidden(esk4_2(X13,X14),X14)|v1_tops_2(X14,X13)|~m1_subset_1(X14,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X13))))|~l1_pre_topc(X13))&(~v3_pre_topc(esk4_2(X13,X14),X13)|v1_tops_2(X14,X13)|~m1_subset_1(X14,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X13))))|~l1_pre_topc(X13))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tops_2])])])])])).
fof(c_0_10, negated_conjecture, (l1_pre_topc(esk1_0)&(m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))&((~v1_tops_2(esk2_0,esk1_0)|~r1_tarski(esk2_0,u1_pre_topc(esk1_0)))&(v1_tops_2(esk2_0,esk1_0)|r1_tarski(esk2_0,u1_pre_topc(esk1_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
cnf(c_0_11, plain, (r1_tarski(X2,X3)|~r2_hidden(esk3_3(X1,X2,X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (r2_hidden(X1,u1_pre_topc(X2))|~v3_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (v3_pre_topc(X3,X2)|~v1_tops_2(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (r1_tarski(X1,u1_pre_topc(X2))|~v3_pre_topc(esk3_3(X3,X1,u1_pre_topc(X2)),X2)|~m1_subset_1(esk3_3(X3,X1,u1_pre_topc(X2)),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(u1_pre_topc(X2),k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_17, negated_conjecture, (v3_pre_topc(X1,esk1_0)|~r2_hidden(X1,esk2_0)|~v1_tops_2(esk2_0,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])])).
cnf(c_0_18, negated_conjecture, (r1_tarski(X1,u1_pre_topc(esk1_0))|~r2_hidden(esk3_3(X2,X1,u1_pre_topc(esk1_0)),esk2_0)|~v1_tops_2(esk2_0,esk1_0)|~m1_subset_1(esk3_3(X2,X1,u1_pre_topc(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_pre_topc(esk1_0),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_15])])).
cnf(c_0_19, plain, (r2_hidden(esk3_3(X1,X2,X3),X2)|r1_tarski(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_20, negated_conjecture, (v1_tops_2(esk2_0,esk1_0)|r1_tarski(esk2_0,u1_pre_topc(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_21, plain, ![X17, X18, X19]:(~l1_pre_topc(X17)|(~m1_subset_1(X18,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X17))))|(~m1_subset_1(X19,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X17))))|(~r1_tarski(X18,X19)|~v1_tops_2(X19,X17)|v1_tops_2(X18,X17))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t18_tops_2])])])).
fof(c_0_22, plain, ![X20]:(~l1_pre_topc(X20)|m1_subset_1(u1_pre_topc(X20),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X20))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_pre_topc])])).
cnf(c_0_23, plain, (v3_pre_topc(X1,X2)|~r2_hidden(X1,u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_24, plain, (r2_hidden(esk4_2(X1,X2),X2)|v1_tops_2(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_25, negated_conjecture, (r1_tarski(esk2_0,u1_pre_topc(esk1_0))|~m1_subset_1(esk3_3(X1,esk2_0,u1_pre_topc(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_pre_topc(esk1_0),k1_zfmisc_1(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])).
cnf(c_0_26, plain, (m1_subset_1(esk3_3(X1,X2,X3),X1)|r1_tarski(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_27, plain, (v1_tops_2(X2,X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~r1_tarski(X2,X3)|~v1_tops_2(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, plain, (v1_tops_2(X2,X1)|~v3_pre_topc(esk4_2(X1,X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_30, plain, (v3_pre_topc(esk4_2(X1,u1_pre_topc(X2)),X2)|v1_tops_2(u1_pre_topc(X2),X1)|~m1_subset_1(esk4_2(X1,u1_pre_topc(X2)),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(u1_pre_topc(X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_31, negated_conjecture, (r1_tarski(esk2_0,u1_pre_topc(esk1_0))|~m1_subset_1(u1_pre_topc(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_14])])).
cnf(c_0_32, plain, (v1_tops_2(X1,X2)|~r1_tarski(X1,u1_pre_topc(X2))|~v1_tops_2(u1_pre_topc(X2),X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_33, plain, (v1_tops_2(u1_pre_topc(X1),X1)|~m1_subset_1(esk4_2(X1,u1_pre_topc(X1)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_28])).
cnf(c_0_34, plain, (m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|v1_tops_2(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_35, negated_conjecture, (~v1_tops_2(esk2_0,esk1_0)|~r1_tarski(esk2_0,u1_pre_topc(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_36, negated_conjecture, (r1_tarski(esk2_0,u1_pre_topc(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_28]), c_0_15])])).
cnf(c_0_37, negated_conjecture, (v1_tops_2(esk2_0,esk1_0)|~v1_tops_2(u1_pre_topc(esk1_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_20]), c_0_14]), c_0_15])])).
cnf(c_0_38, plain, (v1_tops_2(u1_pre_topc(X1),X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_28])).
cnf(c_0_39, negated_conjecture, (~v1_tops_2(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_36])])).
cnf(c_0_40, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_15])]), c_0_39]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 41
# Proof object clause steps            : 28
# Proof object formula steps           : 13
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 19
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 15
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 15
# Processed clauses                    : 52
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 50
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 3
# Generated clauses                    : 32
# ...of the previous two non-trivial   : 26
# Contextual simplify-reflections      : 3
# Paramodulations                      : 32
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 30
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 24
# Current number of unprocessed clauses: 4
# ...number of literals in the above   : 26
# Current number of archived formulas  : 0
# Current number of archived clauses   : 20
# Clause-clause subsumption calls (NU) : 198
# Rec. Clause-clause subsumption calls : 12
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 13
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2494

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28684 pages
