# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t6_yellow_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>(v1_lattice3(k2_yellow_1(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>r1_tarski(k2_xboole_0(X2,X3),k10_lattice3(k2_yellow_1(X1),X2,X3)))))), file('data/MPTP/proofs/330fd3ff3ae04312bb6d105cb0916013.E_input', t6_yellow_1)).
fof(t3_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r3_orders_2(k2_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3))))), file('data/MPTP/proofs/330fd3ff3ae04312bb6d105cb0916013.E_input', t3_yellow_1)).
fof(t1_yellow_1, axiom, ![X1]:(u1_struct_0(k2_yellow_1(X1))=X1&u1_orders_2(k2_yellow_1(X1))=k1_yellow_1(X1)), file('data/MPTP/proofs/330fd3ff3ae04312bb6d105cb0916013.E_input', t1_yellow_1)).
fof(t13_lattice3, axiom, ![X1]:((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k10_lattice3(X1,X2,X3)=k10_lattice3(X1,X3,X2)))), file('data/MPTP/proofs/330fd3ff3ae04312bb6d105cb0916013.E_input', t13_lattice3)).
fof(redefinition_k13_lattice3, axiom, ![X1, X2, X3]:(((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k13_lattice3(X1,X2,X3)=k10_lattice3(X1,X2,X3)), file('data/MPTP/proofs/330fd3ff3ae04312bb6d105cb0916013.E_input', redefinition_k13_lattice3)).
fof(t8_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X3,X2))=>r1_tarski(k2_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/330fd3ff3ae04312bb6d105cb0916013.E_input', t8_xboole_1)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/330fd3ff3ae04312bb6d105cb0916013.E_input', redefinition_r3_orders_2)).
fof(fc5_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&(v3_orders_2(k2_yellow_1(X1))&(v4_orders_2(k2_yellow_1(X1))&v5_orders_2(k2_yellow_1(X1))))), file('data/MPTP/proofs/330fd3ff3ae04312bb6d105cb0916013.E_input', fc5_yellow_1)).
fof(dt_k2_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&l1_orders_2(k2_yellow_1(X1))), file('data/MPTP/proofs/330fd3ff3ae04312bb6d105cb0916013.E_input', dt_k2_yellow_1)).
fof(fc6_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>(~(v2_struct_0(k2_yellow_1(X1)))&v1_orders_2(k2_yellow_1(X1)))), file('data/MPTP/proofs/330fd3ff3ae04312bb6d105cb0916013.E_input', fc6_yellow_1)).
fof(t22_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k13_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))))))), file('data/MPTP/proofs/330fd3ff3ae04312bb6d105cb0916013.E_input', t22_yellow_0)).
fof(dt_k13_lattice3, axiom, ![X1, X2, X3]:(((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k13_lattice3(X1,X2,X3),u1_struct_0(X1))), file('data/MPTP/proofs/330fd3ff3ae04312bb6d105cb0916013.E_input', dt_k13_lattice3)).
fof(c_0_12, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>(v1_lattice3(k2_yellow_1(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>r1_tarski(k2_xboole_0(X2,X3),k10_lattice3(k2_yellow_1(X1),X2,X3))))))), inference(assume_negation,[status(cth)],[t6_yellow_1])).
fof(c_0_13, plain, ![X12, X13, X14]:((~r3_orders_2(k2_yellow_1(X12),X13,X14)|r1_tarski(X13,X14)|~m1_subset_1(X14,u1_struct_0(k2_yellow_1(X12)))|~m1_subset_1(X13,u1_struct_0(k2_yellow_1(X12)))|v1_xboole_0(X12))&(~r1_tarski(X13,X14)|r3_orders_2(k2_yellow_1(X12),X13,X14)|~m1_subset_1(X14,u1_struct_0(k2_yellow_1(X12)))|~m1_subset_1(X13,u1_struct_0(k2_yellow_1(X12)))|v1_xboole_0(X12))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_yellow_1])])])])])).
fof(c_0_14, plain, ![X24]:(u1_struct_0(k2_yellow_1(X24))=X24&u1_orders_2(k2_yellow_1(X24))=k1_yellow_1(X24)), inference(variable_rename,[status(thm)],[t1_yellow_1])).
fof(c_0_15, plain, ![X18, X19, X20]:(~v5_orders_2(X18)|(~v1_lattice3(X18)|~l1_orders_2(X18))|(~m1_subset_1(X19,u1_struct_0(X18))|(~m1_subset_1(X20,u1_struct_0(X18))|k10_lattice3(X18,X19,X20)=k10_lattice3(X18,X20,X19)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_lattice3])])])).
fof(c_0_16, plain, ![X25, X26, X27]:(~v5_orders_2(X25)|(~v1_lattice3(X25)|~l1_orders_2(X25))|(~m1_subset_1(X26,u1_struct_0(X25))|~m1_subset_1(X27,u1_struct_0(X25)))|k13_lattice3(X25,X26,X27)=k10_lattice3(X25,X26,X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k13_lattice3])])).
fof(c_0_17, negated_conjecture, (~v1_xboole_0(esk1_0)&(v1_lattice3(k2_yellow_1(esk1_0))&(m1_subset_1(esk2_0,u1_struct_0(k2_yellow_1(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(k2_yellow_1(esk1_0)))&~r1_tarski(k2_xboole_0(esk2_0,esk3_0),k10_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])).
fof(c_0_18, plain, ![X15, X16, X17]:(~r1_tarski(X15,X16)|~r1_tarski(X17,X16)|r1_tarski(k2_xboole_0(X15,X17),X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_xboole_1])])).
cnf(c_0_19, plain, (r1_tarski(X2,X3)|v1_xboole_0(X1)|~r3_orders_2(k2_yellow_1(X1),X2,X3)|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, plain, (u1_struct_0(k2_yellow_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_21, plain, ![X21, X22, X23]:((~r3_orders_2(X21,X22,X23)|r1_orders_2(X21,X22,X23)|(v2_struct_0(X21)|(~v3_orders_2(X21)|~l1_orders_2(X21))|(~m1_subset_1(X22,u1_struct_0(X21))|~m1_subset_1(X23,u1_struct_0(X21)))))&(~r1_orders_2(X21,X22,X23)|r3_orders_2(X21,X22,X23)|(v2_struct_0(X21)|(~v3_orders_2(X21)|~l1_orders_2(X21))|(~m1_subset_1(X22,u1_struct_0(X21))|~m1_subset_1(X23,u1_struct_0(X21)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
fof(c_0_22, plain, ![X34]:(v1_orders_2(k2_yellow_1(X34))&(v3_orders_2(k2_yellow_1(X34))&(v4_orders_2(k2_yellow_1(X34))&v5_orders_2(k2_yellow_1(X34))))), inference(variable_rename,[status(thm)],[fc5_yellow_1])).
fof(c_0_23, plain, ![X35]:(v1_orders_2(k2_yellow_1(X35))&l1_orders_2(k2_yellow_1(X35))), inference(variable_rename,[status(thm)],[dt_k2_yellow_1])).
fof(c_0_24, plain, ![X36]:((~v2_struct_0(k2_yellow_1(X36))|v1_xboole_0(X36))&(v1_orders_2(k2_yellow_1(X36))|v1_xboole_0(X36))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc6_yellow_1])])])])).
fof(c_0_25, plain, ![X28, X29, X30, X31, X32]:(((r1_orders_2(X28,X29,X31)|X31!=k13_lattice3(X28,X29,X30)|~m1_subset_1(X31,u1_struct_0(X28))|~m1_subset_1(X30,u1_struct_0(X28))|~m1_subset_1(X29,u1_struct_0(X28))|(~v5_orders_2(X28)|(~v1_lattice3(X28)|~l1_orders_2(X28))))&((r1_orders_2(X28,X30,X31)|X31!=k13_lattice3(X28,X29,X30)|~m1_subset_1(X31,u1_struct_0(X28))|~m1_subset_1(X30,u1_struct_0(X28))|~m1_subset_1(X29,u1_struct_0(X28))|(~v5_orders_2(X28)|(~v1_lattice3(X28)|~l1_orders_2(X28))))&(~m1_subset_1(X32,u1_struct_0(X28))|(~r1_orders_2(X28,X29,X32)|~r1_orders_2(X28,X30,X32)|r1_orders_2(X28,X31,X32))|X31!=k13_lattice3(X28,X29,X30)|~m1_subset_1(X31,u1_struct_0(X28))|~m1_subset_1(X30,u1_struct_0(X28))|~m1_subset_1(X29,u1_struct_0(X28))|(~v5_orders_2(X28)|(~v1_lattice3(X28)|~l1_orders_2(X28))))))&((m1_subset_1(esk4_4(X28,X29,X30,X31),u1_struct_0(X28))|~r1_orders_2(X28,X30,X31)|~r1_orders_2(X28,X29,X31)|X31=k13_lattice3(X28,X29,X30)|~m1_subset_1(X31,u1_struct_0(X28))|~m1_subset_1(X30,u1_struct_0(X28))|~m1_subset_1(X29,u1_struct_0(X28))|(~v5_orders_2(X28)|(~v1_lattice3(X28)|~l1_orders_2(X28))))&(((r1_orders_2(X28,X29,esk4_4(X28,X29,X30,X31))|~r1_orders_2(X28,X30,X31)|~r1_orders_2(X28,X29,X31)|X31=k13_lattice3(X28,X29,X30)|~m1_subset_1(X31,u1_struct_0(X28))|~m1_subset_1(X30,u1_struct_0(X28))|~m1_subset_1(X29,u1_struct_0(X28))|(~v5_orders_2(X28)|(~v1_lattice3(X28)|~l1_orders_2(X28))))&(r1_orders_2(X28,X30,esk4_4(X28,X29,X30,X31))|~r1_orders_2(X28,X30,X31)|~r1_orders_2(X28,X29,X31)|X31=k13_lattice3(X28,X29,X30)|~m1_subset_1(X31,u1_struct_0(X28))|~m1_subset_1(X30,u1_struct_0(X28))|~m1_subset_1(X29,u1_struct_0(X28))|(~v5_orders_2(X28)|(~v1_lattice3(X28)|~l1_orders_2(X28)))))&(~r1_orders_2(X28,X31,esk4_4(X28,X29,X30,X31))|~r1_orders_2(X28,X30,X31)|~r1_orders_2(X28,X29,X31)|X31=k13_lattice3(X28,X29,X30)|~m1_subset_1(X31,u1_struct_0(X28))|~m1_subset_1(X30,u1_struct_0(X28))|~m1_subset_1(X29,u1_struct_0(X28))|(~v5_orders_2(X28)|(~v1_lattice3(X28)|~l1_orders_2(X28))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t22_yellow_0])])])])])).
fof(c_0_26, plain, ![X9, X10, X11]:(~v5_orders_2(X9)|(~v1_lattice3(X9)|~l1_orders_2(X9))|(~m1_subset_1(X10,u1_struct_0(X9))|~m1_subset_1(X11,u1_struct_0(X9)))|m1_subset_1(k13_lattice3(X9,X10,X11),u1_struct_0(X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k13_lattice3])])).
cnf(c_0_27, plain, (k10_lattice3(X1,X2,X3)=k10_lattice3(X1,X3,X2)|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_28, plain, (k13_lattice3(X1,X2,X3)=k10_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_29, negated_conjecture, (~r1_tarski(k2_xboole_0(esk2_0,esk3_0),k10_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_30, plain, (r1_tarski(k2_xboole_0(X1,X3),X2)|~r1_tarski(X1,X2)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k2_yellow_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_32, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(k2_yellow_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_33, plain, (r1_tarski(X1,X2)|v1_xboole_0(X3)|~r3_orders_2(k2_yellow_1(X3),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_20]), c_0_20])).
cnf(c_0_34, plain, (r3_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_35, plain, (v3_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_36, plain, (l1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_37, plain, (v1_xboole_0(X1)|~v2_struct_0(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_38, plain, (r1_orders_2(X1,X2,X3)|X3!=k13_lattice3(X1,X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_39, plain, (m1_subset_1(k13_lattice3(X1,X2,X3),u1_struct_0(X1))|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_40, plain, (v5_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_41, plain, (k10_lattice3(X1,X2,X3)=k13_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_lattice3(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_42, negated_conjecture, (~r1_tarski(esk3_0,k10_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0))|~r1_tarski(esk2_0,k10_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_43, negated_conjecture, (m1_subset_1(esk3_0,esk1_0)), inference(rw,[status(thm)],[c_0_31, c_0_20])).
cnf(c_0_44, negated_conjecture, (m1_subset_1(esk2_0,esk1_0)), inference(rw,[status(thm)],[c_0_32, c_0_20])).
cnf(c_0_45, negated_conjecture, (v1_lattice3(k2_yellow_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_46, plain, (r1_tarski(X1,X2)|v1_xboole_0(X3)|~r1_orders_2(k2_yellow_1(X3),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35]), c_0_36]), c_0_20]), c_0_20])]), c_0_37])).
cnf(c_0_47, plain, (r1_orders_2(X1,X2,k13_lattice3(X1,X3,X2))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_lattice3(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_38]), c_0_39])).
cnf(c_0_48, plain, (m1_subset_1(k13_lattice3(k2_yellow_1(X1),X2,X3),X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v1_lattice3(k2_yellow_1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_20]), c_0_36]), c_0_40])])).
cnf(c_0_49, plain, (k13_lattice3(X1,X2,X3)=k13_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_lattice3(X1)), inference(spm,[status(thm)],[c_0_28, c_0_41])).
cnf(c_0_50, negated_conjecture, (~r1_tarski(esk3_0,k13_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0))|~r1_tarski(esk2_0,k13_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_28]), c_0_36]), c_0_40]), c_0_20]), c_0_43]), c_0_20]), c_0_44]), c_0_45])])).
cnf(c_0_51, plain, (r1_tarski(X1,k13_lattice3(k2_yellow_1(X2),X3,X1))|v1_xboole_0(X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X3,X2)|~v1_lattice3(k2_yellow_1(X2))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_36]), c_0_40]), c_0_20]), c_0_20])]), c_0_48])).
cnf(c_0_52, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_53, plain, (r1_orders_2(X1,X2,k13_lattice3(X1,X2,X3))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_lattice3(X1)), inference(spm,[status(thm)],[c_0_47, c_0_49])).
cnf(c_0_54, negated_conjecture, (~r1_tarski(esk2_0,k13_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_43]), c_0_44]), c_0_45])]), c_0_52])).
cnf(c_0_55, plain, (r1_tarski(X1,k13_lattice3(k2_yellow_1(X2),X1,X3))|v1_xboole_0(X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X3,X2)|~v1_lattice3(k2_yellow_1(X2))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_53]), c_0_36]), c_0_40]), c_0_20]), c_0_20])]), c_0_48])).
cnf(c_0_56, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_44]), c_0_43]), c_0_45])]), c_0_52]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 57
# Proof object clause steps            : 32
# Proof object formula steps           : 25
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 45
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 29
# Processed clauses                    : 57
# ...of these trivial                  : 2
# ...subsumed                          : 17
# ...remaining for further processing  : 38
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 44
# ...of the previous two non-trivial   : 38
# Contextual simplify-reflections      : 5
# Paramodulations                      : 43
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 37
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 23
# Current number of unprocessed clauses: 9
# ...number of literals in the above   : 78
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1
# Clause-clause subsumption calls (NU) : 137
# Rec. Clause-clause subsumption calls : 33
# Non-unit clause-clause subsumptions  : 20
# Unit Clause-clause subsumption calls : 12
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3692

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28728 pages
