# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t87_orders_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_orders_1(X2,k1_orders_1(u1_struct_0(X1)))=>~(k3_tarski(k4_orders_2(X1,X2))=k1_xboole_0))), file('data/MPTP/proofs/3321ab2c56b246be9cd96fe64a45c74a.E_input', t87_orders_2)).
fof(l49_zfmisc_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>r1_tarski(X1,k3_tarski(X2))), file('data/MPTP/proofs/3321ab2c56b246be9cd96fe64a45c74a.E_input', l49_zfmisc_1)).
fof(t34_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4, X5, X6]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k4_mcart_1(X3,X4,X5,X6))))))), file('data/MPTP/proofs/3321ab2c56b246be9cd96fe64a45c74a.E_input', t34_mcart_1)).
fof(d16_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_orders_1(X2,k1_orders_1(u1_struct_0(X1)))=>![X3]:((v6_orders_2(X3,X1)&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))=>(m2_orders_2(X3,X1,X2)<=>(~(X3=k1_xboole_0)&(r2_wellord1(u1_orders_2(X1),X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>k1_funct_1(X2,k1_orders_2(X1,k3_orders_2(X1,X3,X4)))=X4)))))))), file('data/MPTP/proofs/3321ab2c56b246be9cd96fe64a45c74a.E_input', d16_orders_2)).
fof(dt_m2_orders_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))&m1_orders_1(X2,k1_orders_1(u1_struct_0(X1))))=>![X3]:(m2_orders_2(X3,X1,X2)=>(v6_orders_2(X3,X1)&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/3321ab2c56b246be9cd96fe64a45c74a.E_input', dt_m2_orders_2)).
fof(t3_xboole_1, axiom, ![X1]:(r1_tarski(X1,k1_xboole_0)=>X1=k1_xboole_0), file('data/MPTP/proofs/3321ab2c56b246be9cd96fe64a45c74a.E_input', t3_xboole_1)).
fof(d17_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_orders_1(X2,k1_orders_1(u1_struct_0(X1)))=>![X3]:(X3=k4_orders_2(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>m2_orders_2(X4,X1,X2))))), file('data/MPTP/proofs/3321ab2c56b246be9cd96fe64a45c74a.E_input', d17_orders_2)).
fof(fc9_orders_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))&m1_orders_1(X2,k1_orders_1(u1_struct_0(X1))))=>~(v1_xboole_0(k4_orders_2(X1,X2)))), file('data/MPTP/proofs/3321ab2c56b246be9cd96fe64a45c74a.E_input', fc9_orders_2)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/3321ab2c56b246be9cd96fe64a45c74a.E_input', fc1_xboole_0)).
fof(c_0_9, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_orders_1(X2,k1_orders_1(u1_struct_0(X1)))=>~(k3_tarski(k4_orders_2(X1,X2))=k1_xboole_0)))), inference(assume_negation,[status(cth)],[t87_orders_2])).
fof(c_0_10, plain, ![X9, X10]:(~r2_hidden(X9,X10)|r1_tarski(X9,k3_tarski(X10))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l49_zfmisc_1])])).
fof(c_0_11, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&(m1_orders_1(esk2_0,k1_orders_1(u1_struct_0(esk1_0)))&k3_tarski(k4_orders_2(esk1_0,esk2_0))=k1_xboole_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])).
cnf(c_0_12, plain, (r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_13, negated_conjecture, (k3_tarski(k4_orders_2(esk1_0,esk2_0))=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_14, plain, ![X23, X25, X26, X27, X28]:((r2_hidden(esk4_1(X23),X23)|X23=k1_xboole_0)&((~r2_hidden(X25,X23)|esk4_1(X23)!=k4_mcart_1(X25,X26,X27,X28)|X23=k1_xboole_0)&(~r2_hidden(X26,X23)|esk4_1(X23)!=k4_mcart_1(X25,X26,X27,X28)|X23=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_mcart_1])])])])])])])).
fof(c_0_15, plain, ![X29, X30, X31, X32]:(((X31!=k1_xboole_0|~m2_orders_2(X31,X29,X30)|(~v6_orders_2(X31,X29)|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X29))))|~m1_orders_1(X30,k1_orders_1(u1_struct_0(X29)))|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))))&((r2_wellord1(u1_orders_2(X29),X31)|~m2_orders_2(X31,X29,X30)|(~v6_orders_2(X31,X29)|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X29))))|~m1_orders_1(X30,k1_orders_1(u1_struct_0(X29)))|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))))&(~m1_subset_1(X32,u1_struct_0(X29))|(~r2_hidden(X32,X31)|k1_funct_1(X30,k1_orders_2(X29,k3_orders_2(X29,X31,X32)))=X32)|~m2_orders_2(X31,X29,X30)|(~v6_orders_2(X31,X29)|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X29))))|~m1_orders_1(X30,k1_orders_1(u1_struct_0(X29)))|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))))))&((m1_subset_1(esk5_3(X29,X30,X31),u1_struct_0(X29))|~r2_wellord1(u1_orders_2(X29),X31)|X31=k1_xboole_0|m2_orders_2(X31,X29,X30)|(~v6_orders_2(X31,X29)|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X29))))|~m1_orders_1(X30,k1_orders_1(u1_struct_0(X29)))|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))))&((r2_hidden(esk5_3(X29,X30,X31),X31)|~r2_wellord1(u1_orders_2(X29),X31)|X31=k1_xboole_0|m2_orders_2(X31,X29,X30)|(~v6_orders_2(X31,X29)|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X29))))|~m1_orders_1(X30,k1_orders_1(u1_struct_0(X29)))|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))))&(k1_funct_1(X30,k1_orders_2(X29,k3_orders_2(X29,X31,esk5_3(X29,X30,X31))))!=esk5_3(X29,X30,X31)|~r2_wellord1(u1_orders_2(X29),X31)|X31=k1_xboole_0|m2_orders_2(X31,X29,X30)|(~v6_orders_2(X31,X29)|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X29))))|~m1_orders_1(X30,k1_orders_1(u1_struct_0(X29)))|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_orders_2])])])])])])).
fof(c_0_16, plain, ![X13, X14, X15]:((v6_orders_2(X15,X13)|~m2_orders_2(X15,X13,X14)|(v2_struct_0(X13)|(~v3_orders_2(X13)|(~v4_orders_2(X13)|(~v5_orders_2(X13)|~l1_orders_2(X13))))|~m1_orders_1(X14,k1_orders_1(u1_struct_0(X13)))))&(m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X13)))|~m2_orders_2(X15,X13,X14)|(v2_struct_0(X13)|(~v3_orders_2(X13)|(~v4_orders_2(X13)|(~v5_orders_2(X13)|~l1_orders_2(X13))))|~m1_orders_1(X14,k1_orders_1(u1_struct_0(X13)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_orders_2])])])])])).
fof(c_0_17, plain, ![X34]:(~r1_tarski(X34,k1_xboole_0)|X34=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_xboole_1])])).
cnf(c_0_18, negated_conjecture, (r1_tarski(X1,k1_xboole_0)|~r2_hidden(X1,k4_orders_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_19, plain, (r2_hidden(esk4_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (v2_struct_0(X2)|X1!=k1_xboole_0|~m2_orders_2(X1,X2,X3)|~v6_orders_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_orders_1(X3,k1_orders_1(u1_struct_0(X2)))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (v6_orders_2(X1,X2)|v2_struct_0(X2)|~m2_orders_2(X1,X2,X3)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_orders_1(X3,k1_orders_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|~m2_orders_2(X1,X2,X3)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_orders_1(X3,k1_orders_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_23, plain, ![X16, X17, X18, X19, X20, X21]:(((~r2_hidden(X19,X18)|m2_orders_2(X19,X16,X17)|X18!=k4_orders_2(X16,X17)|~m1_orders_1(X17,k1_orders_1(u1_struct_0(X16)))|(v2_struct_0(X16)|(~v3_orders_2(X16)|(~v4_orders_2(X16)|(~v5_orders_2(X16)|~l1_orders_2(X16))))))&(~m2_orders_2(X20,X16,X17)|r2_hidden(X20,X18)|X18!=k4_orders_2(X16,X17)|~m1_orders_1(X17,k1_orders_1(u1_struct_0(X16)))|(v2_struct_0(X16)|(~v3_orders_2(X16)|(~v4_orders_2(X16)|(~v5_orders_2(X16)|~l1_orders_2(X16)))))))&((~r2_hidden(esk3_3(X16,X17,X21),X21)|~m2_orders_2(esk3_3(X16,X17,X21),X16,X17)|X21=k4_orders_2(X16,X17)|~m1_orders_1(X17,k1_orders_1(u1_struct_0(X16)))|(v2_struct_0(X16)|(~v3_orders_2(X16)|(~v4_orders_2(X16)|(~v5_orders_2(X16)|~l1_orders_2(X16))))))&(r2_hidden(esk3_3(X16,X17,X21),X21)|m2_orders_2(esk3_3(X16,X17,X21),X16,X17)|X21=k4_orders_2(X16,X17)|~m1_orders_1(X17,k1_orders_1(u1_struct_0(X16)))|(v2_struct_0(X16)|(~v3_orders_2(X16)|(~v4_orders_2(X16)|(~v5_orders_2(X16)|~l1_orders_2(X16)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_orders_2])])])])])])])).
cnf(c_0_24, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, negated_conjecture, (k4_orders_2(esk1_0,esk2_0)=k1_xboole_0|r1_tarski(esk4_1(k4_orders_2(esk1_0,esk2_0)),k1_xboole_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_26, plain, (v2_struct_0(X1)|X2!=k1_xboole_0|~m2_orders_2(X2,X1,X3)|~m1_orders_1(X3,k1_orders_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_20, c_0_21]), c_0_22])).
cnf(c_0_27, plain, (m2_orders_2(X1,X3,X4)|v2_struct_0(X3)|~r2_hidden(X1,X2)|X2!=k4_orders_2(X3,X4)|~m1_orders_1(X4,k1_orders_1(u1_struct_0(X3)))|~v3_orders_2(X3)|~v4_orders_2(X3)|~v5_orders_2(X3)|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28, negated_conjecture, (esk4_1(k4_orders_2(esk1_0,esk2_0))=k1_xboole_0|k4_orders_2(esk1_0,esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29, plain, (v2_struct_0(X1)|~m2_orders_2(k1_xboole_0,X1,X2)|~m1_orders_1(X2,k1_orders_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(er,[status(thm)],[c_0_26])).
cnf(c_0_30, negated_conjecture, (m1_orders_1(esk2_0,k1_orders_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_31, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_32, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_33, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_34, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_35, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_36, plain, ![X11, X12]:(v2_struct_0(X11)|(~v3_orders_2(X11)|(~v4_orders_2(X11)|(~v5_orders_2(X11)|~l1_orders_2(X11))))|~m1_orders_1(X12,k1_orders_1(u1_struct_0(X11)))|~v1_xboole_0(k4_orders_2(X11,X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc9_orders_2])])])).
cnf(c_0_37, plain, (m2_orders_2(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,k4_orders_2(X2,X3))|~m1_orders_1(X3,k1_orders_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_38, negated_conjecture, (k4_orders_2(esk1_0,esk2_0)=k1_xboole_0|r2_hidden(k1_xboole_0,k4_orders_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_19, c_0_28])).
cnf(c_0_39, negated_conjecture, (~m2_orders_2(k1_xboole_0,esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_40, plain, (v2_struct_0(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_orders_1(X2,k1_orders_1(u1_struct_0(X1)))|~v1_xboole_0(k4_orders_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_41, negated_conjecture, (k4_orders_2(esk1_0,esk2_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34])]), c_0_39]), c_0_35])).
cnf(c_0_42, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_43, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34])]), c_0_35]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 44
# Proof object clause steps            : 26
# Proof object formula steps           : 18
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 24
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 26
# Processed clauses                    : 76
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 75
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 11
# Generated clauses                    : 32
# ...of the previous two non-trivial   : 31
# Contextual simplify-reflections      : 12
# Paramodulations                      : 29
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 35
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 26
# Current number of unprocessed clauses: 2
# ...number of literals in the above   : 17
# Current number of archived formulas  : 0
# Current number of archived clauses   : 40
# Clause-clause subsumption calls (NU) : 884
# Rec. Clause-clause subsumption calls : 81
# Non-unit clause-clause subsumptions  : 16
# Unit Clause-clause subsumption calls : 28
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4201

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28616 pages
