# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/339b6ea2b03f4675a5f0d3bc9b402327.E_input', d4_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/339b6ea2b03f4675a5f0d3bc9b402327.E_input', d3_tarski)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/339b6ea2b03f4675a5f0d3bc9b402327.E_input', t113_zfmisc_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/339b6ea2b03f4675a5f0d3bc9b402327.E_input', t3_boole)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/339b6ea2b03f4675a5f0d3bc9b402327.E_input', l32_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/339b6ea2b03f4675a5f0d3bc9b402327.E_input', t28_xboole_1)).
fof(t123_zfmisc_1, axiom, ![X1, X2, X3, X4]:k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)), file('data/MPTP/proofs/339b6ea2b03f4675a5f0d3bc9b402327.E_input', t123_zfmisc_1)).
fof(t118_zfmisc_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>(r1_tarski(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&r1_tarski(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2)))), file('data/MPTP/proofs/339b6ea2b03f4675a5f0d3bc9b402327.E_input', t118_zfmisc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/339b6ea2b03f4675a5f0d3bc9b402327.E_input', d10_xboole_0)).
fof(t138_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=>(k2_zfmisc_1(X1,X2)=k1_xboole_0|(r1_tarski(X1,X3)&r1_tarski(X2,X4)))), file('data/MPTP/proofs/339b6ea2b03f4675a5f0d3bc9b402327.E_input', t138_zfmisc_1)).
fof(t117_zfmisc_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&((r1_tarski(k2_zfmisc_1(X2,X1),k2_zfmisc_1(X3,X1))|r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))&~(r1_tarski(X2,X3))))), file('data/MPTP/proofs/339b6ea2b03f4675a5f0d3bc9b402327.E_input', t117_zfmisc_1)).
fof(c_0_11, plain, ![X24, X25, X26, X27, X28, X29, X30, X31]:((((r2_hidden(X27,X24)|~r2_hidden(X27,X26)|X26!=k3_xboole_0(X24,X25))&(r2_hidden(X27,X25)|~r2_hidden(X27,X26)|X26!=k3_xboole_0(X24,X25)))&(~r2_hidden(X28,X24)|~r2_hidden(X28,X25)|r2_hidden(X28,X26)|X26!=k3_xboole_0(X24,X25)))&((~r2_hidden(esk5_3(X29,X30,X31),X31)|(~r2_hidden(esk5_3(X29,X30,X31),X29)|~r2_hidden(esk5_3(X29,X30,X31),X30))|X31=k3_xboole_0(X29,X30))&((r2_hidden(esk5_3(X29,X30,X31),X29)|r2_hidden(esk5_3(X29,X30,X31),X31)|X31=k3_xboole_0(X29,X30))&(r2_hidden(esk5_3(X29,X30,X31),X30)|r2_hidden(esk5_3(X29,X30,X31),X31)|X31=k3_xboole_0(X29,X30))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_12, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_13, plain, ![X33, X34, X35, X36, X37]:((~r1_tarski(X33,X34)|(~r2_hidden(X35,X33)|r2_hidden(X35,X34)))&((r2_hidden(esk6_2(X36,X37),X36)|r1_tarski(X36,X37))&(~r2_hidden(esk6_2(X36,X37),X37)|r1_tarski(X36,X37)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_14, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_12])).
cnf(c_0_15, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_16, plain, (r2_hidden(esk6_2(k3_xboole_0(X1,X2),X3),X2)|r1_tarski(k3_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
fof(c_0_17, plain, ![X20, X21]:((k2_zfmisc_1(X20,X21)!=k1_xboole_0|(X20=k1_xboole_0|X21=k1_xboole_0))&((X20!=k1_xboole_0|k2_zfmisc_1(X20,X21)=k1_xboole_0)&(X21!=k1_xboole_0|k2_zfmisc_1(X20,X21)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_18, plain, ![X9]:k4_xboole_0(X9,k1_xboole_0)=X9, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_19, plain, ![X22, X23]:((k4_xboole_0(X22,X23)!=k1_xboole_0|r1_tarski(X22,X23))&(~r1_tarski(X22,X23)|k4_xboole_0(X22,X23)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_20, plain, (r1_tarski(X1,X2)|~r2_hidden(esk6_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (r2_hidden(esk6_2(k3_xboole_0(X1,k3_xboole_0(X2,X3)),X4),X3)|r1_tarski(k3_xboole_0(X1,k3_xboole_0(X2,X3)),X4)), inference(spm,[status(thm)],[c_0_14, c_0_16])).
fof(c_0_22, plain, ![X15, X16]:(~r1_tarski(X15,X16)|k3_xboole_0(X15,X16)=X15), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_23, plain, ![X39, X40, X41, X42]:k2_zfmisc_1(k3_xboole_0(X39,X40),k3_xboole_0(X41,X42))=k3_xboole_0(k2_zfmisc_1(X39,X41),k2_zfmisc_1(X40,X42)), inference(variable_rename,[status(thm)],[t123_zfmisc_1])).
cnf(c_0_24, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_25, plain, ![X10, X11, X12]:((r1_tarski(k2_zfmisc_1(X10,X12),k2_zfmisc_1(X11,X12))|~r1_tarski(X10,X11))&(r1_tarski(k2_zfmisc_1(X12,X10),k2_zfmisc_1(X12,X11))|~r1_tarski(X10,X11))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t118_zfmisc_1])])])).
cnf(c_0_26, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (r1_tarski(k3_xboole_0(X1,k3_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_30, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_24])).
fof(c_0_33, plain, ![X13, X14]:(((r1_tarski(X13,X14)|X13!=X14)&(r1_tarski(X14,X13)|X13!=X14))&(~r1_tarski(X13,X14)|~r1_tarski(X14,X13)|X13=X14)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_34, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_26])).
fof(c_0_36, negated_conjecture, ~(![X1, X2, X3, X4]:(r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=>(k2_zfmisc_1(X1,X2)=k1_xboole_0|(r1_tarski(X1,X3)&r1_tarski(X2,X4))))), inference(assume_negation,[status(cth)],[t138_zfmisc_1])).
cnf(c_0_37, plain, (k1_xboole_0=X1|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_38, plain, (r1_tarski(k3_xboole_0(X1,X2),X3)|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_39, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_40, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k1_xboole_0)=k2_zfmisc_1(k3_xboole_0(X1,X3),k3_xboole_0(X2,k1_xboole_0))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_41, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k1_xboole_0)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_43, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k2_zfmisc_1(X1,X3)|~r1_tarski(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(spm,[status(thm)],[c_0_31, c_0_30])).
fof(c_0_44, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))&(k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0&(~r1_tarski(esk1_0,esk3_0)|~r1_tarski(esk2_0,esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_36])])])).
cnf(c_0_45, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_46, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_47, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_39])).
cnf(c_0_48, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k1_xboole_0)=k2_zfmisc_1(X1,k3_xboole_0(X2,k1_xboole_0))|~r1_tarski(X1,X3)), inference(spm,[status(thm)],[c_0_40, c_0_30])).
cnf(c_0_49, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_50, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k1_xboole_0)|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))|~r1_tarski(k3_xboole_0(X1,X3),k1_xboole_0)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_51, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_52, plain, (k1_xboole_0=X1|~r1_tarski(X2,k1_xboole_0)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_30, c_0_45])).
cnf(c_0_53, plain, (r1_tarski(X1,k1_xboole_0)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_46, c_0_27])).
cnf(c_0_54, plain, (r2_hidden(esk6_2(k3_xboole_0(X1,X2),X3),X1)|r1_tarski(k3_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_47, c_0_15])).
cnf(c_0_55, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k1_xboole_0)=k2_zfmisc_1(X1,k3_xboole_0(X2,k1_xboole_0))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_56, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k1_xboole_0)|~r1_tarski(k3_xboole_0(esk1_0,esk3_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_57, plain, (k1_xboole_0=X1|X2!=k1_xboole_0|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_58, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_20, c_0_54])).
cnf(c_0_59, plain, (r2_hidden(X1,k1_xboole_0)|X2!=k2_zfmisc_1(X3,k3_xboole_0(X4,k1_xboole_0))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_12, c_0_55])).
cnf(c_0_60, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k1_xboole_0)|k3_xboole_0(esk1_0,esk3_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_56, c_0_53])).
cnf(c_0_61, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_62, plain, (r2_hidden(X1,k1_xboole_0)|X2!=k2_zfmisc_1(X3,X4)|~r2_hidden(X1,X2)|~r1_tarski(X4,k1_xboole_0)), inference(spm,[status(thm)],[c_0_59, c_0_30])).
fof(c_0_63, plain, ![X17, X18, X19]:((~r1_tarski(k2_zfmisc_1(X18,X17),k2_zfmisc_1(X19,X17))|r1_tarski(X18,X19)|X17=k1_xboole_0)&(~r1_tarski(k2_zfmisc_1(X17,X18),k2_zfmisc_1(X17,X19))|r1_tarski(X18,X19)|X17=k1_xboole_0)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t117_zfmisc_1])])])])).
cnf(c_0_64, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,k3_xboole_0(X2,X4)))|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X5,X4))|~r1_tarski(k3_xboole_0(X1,X5),X3)), inference(spm,[status(thm)],[c_0_34, c_0_43])).
cnf(c_0_65, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k1_xboole_0)|esk1_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_66, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_67, plain, (r2_hidden(X1,k1_xboole_0)|X2!=k2_zfmisc_1(X3,X4)|X4!=k1_xboole_0|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_62, c_0_53])).
cnf(c_0_68, plain, (r1_tarski(X2,X3)|X1=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_69, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(X1,k3_xboole_0(esk2_0,esk4_0)))|~r1_tarski(k3_xboole_0(esk1_0,esk3_0),X1)), inference(spm,[status(thm)],[c_0_64, c_0_51])).
cnf(c_0_70, negated_conjecture, (esk1_0!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_65]), c_0_66])).
cnf(c_0_71, plain, (r2_hidden(X1,k1_xboole_0)|X2!=k1_xboole_0|~r2_hidden(X1,k2_zfmisc_1(X3,X2))), inference(er,[status(thm)],[c_0_67])).
cnf(c_0_72, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_73, negated_conjecture, (r1_tarski(esk2_0,k3_xboole_0(esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_58])]), c_0_70])).
cnf(c_0_74, plain, (r2_hidden(esk6_2(k2_zfmisc_1(X1,X2),X3),k1_xboole_0)|r1_tarski(k2_zfmisc_1(X1,X2),X3)|X2!=k1_xboole_0), inference(spm,[status(thm)],[c_0_71, c_0_15])).
cnf(c_0_75, negated_conjecture, (k3_xboole_0(esk2_0,esk4_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_58])])).
cnf(c_0_76, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_20, c_0_16])).
cnf(c_0_77, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k1_xboole_0)|X2!=k1_xboole_0), inference(spm,[status(thm)],[c_0_20, c_0_74])).
cnf(c_0_78, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(X1,esk2_0))|~r1_tarski(k3_xboole_0(esk1_0,esk3_0),X1)), inference(rw,[status(thm)],[c_0_69, c_0_75])).
cnf(c_0_79, negated_conjecture, (~r1_tarski(esk1_0,esk3_0)|~r1_tarski(esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_80, negated_conjecture, (r1_tarski(esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_76, c_0_75])).
cnf(c_0_81, plain, (k1_xboole_0=X1|X2!=k1_xboole_0|~r1_tarski(X1,k2_zfmisc_1(X3,X2))), inference(spm,[status(thm)],[c_0_52, c_0_77])).
cnf(c_0_82, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,esk2_0))), inference(spm,[status(thm)],[c_0_78, c_0_76])).
cnf(c_0_83, plain, (r1_tarski(X1,X3)|X2=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_84, negated_conjecture, (~r1_tarski(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_80])])).
cnf(c_0_85, negated_conjecture, (esk2_0!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_66])).
cnf(c_0_86, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_82]), c_0_84]), c_0_85]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 87
# Proof object clause steps            : 64
# Proof object formula steps           : 23
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 42
# Proof object simplifying inferences  : 13
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 12862
# ...of these trivial                  : 202
# ...subsumed                          : 9784
# ...remaining for further processing  : 2875
# Other redundant clauses eliminated   : 11
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 810
# Backward-rewritten                   : 202
# Generated clauses                    : 184717
# ...of the previous two non-trivial   : 176851
# Contextual simplify-reflections      : 293
# Paramodulations                      : 184424
# Factorizations                       : 210
# Equation resolutions                 : 83
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1835
#    Positive orientable unit clauses  : 33
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 11
#    Non-unit-clauses                  : 1791
# Current number of unprocessed clauses: 159920
# ...number of literals in the above   : 802955
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1038
# Clause-clause subsumption calls (NU) : 628646
# Rec. Clause-clause subsumption calls : 155688
# Non-unit clause-clause subsumptions  : 7770
# Unit Clause-clause subsumption calls : 8675
# Rewrite failures with RHS unbound    : 88
# BW rewrite match attempts            : 189
# BW rewrite match successes           : 59
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3900389

# -------------------------------------------------
# User time                : 2.699 s
# System time              : 0.056 s
# Total time               : 2.756 s
# Maximum resident set size: 127424 pages
