# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t67_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((~(v2_struct_0(X2))&l1_orders_2(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v23_waybel_0(X3,X1,X2)=>((v1_funct_1(k2_funct_1(X3))&(v1_funct_2(k2_funct_1(X3),u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(k2_funct_1(X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))&k2_relat_1(k2_funct_1(X3))=u1_struct_0(X1)))))), file('data/MPTP/proofs/33af5711eb9249929d9d6e789ab7fbac.E_input', t67_waybel_0)).
fof(d38_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(~((~(v2_struct_0(X1))&(~(v2_struct_0(X2))&~((v23_waybel_0(X3,X1,X2)<=>(v2_funct_1(X3)&(v5_orders_3(X3,X1,X2)&?[X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))&(X4=k2_funct_1(X3)&v5_orders_3(X4,X2,X1))))))))))&(~((~(v2_struct_0(X1))&~(v2_struct_0(X2))))=>(v23_waybel_0(X3,X1,X2)<=>(v2_struct_0(X1)&v2_struct_0(X2)))))))), file('data/MPTP/proofs/33af5711eb9249929d9d6e789ab7fbac.E_input', d38_waybel_0)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/33af5711eb9249929d9d6e789ab7fbac.E_input', cc1_relset_1)).
fof(fc6_funct_1, axiom, ![X1]:((v1_relat_1(X1)&(v1_funct_1(X1)&v2_funct_1(X1)))=>(v1_relat_1(k2_funct_1(X1))&(v1_funct_1(k2_funct_1(X1))&v2_funct_1(k2_funct_1(X1))))), file('data/MPTP/proofs/33af5711eb9249929d9d6e789ab7fbac.E_input', fc6_funct_1)).
fof(t66_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((~(v2_struct_0(X2))&l1_orders_2(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v23_waybel_0(X3,X1,X2)<=>(v2_funct_1(X3)&(k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=u1_struct_0(X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r1_orders_2(X1,X4,X5)<=>r1_orders_2(X2,k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,X4),k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,X5))))))))))), file('data/MPTP/proofs/33af5711eb9249929d9d6e789ab7fbac.E_input', t66_waybel_0)).
fof(t65_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)=>k2_funct_1(k2_funct_1(X1))=X1)), file('data/MPTP/proofs/33af5711eb9249929d9d6e789ab7fbac.E_input', t65_funct_1)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/33af5711eb9249929d9d6e789ab7fbac.E_input', redefinition_k2_relset_1)).
fof(c_0_7, plain, ![X1, X3, X2]:(epred1_3(X2,X3,X1)<=>(~((~(v2_struct_0(X1))&(~(v2_struct_0(X2))&~((v23_waybel_0(X3,X1,X2)<=>(v2_funct_1(X3)&(v5_orders_3(X3,X1,X2)&?[X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))&(X4=k2_funct_1(X3)&v5_orders_3(X4,X2,X1))))))))))&(~((~(v2_struct_0(X1))&~(v2_struct_0(X2))))=>(v23_waybel_0(X3,X1,X2)<=>(v2_struct_0(X1)&v2_struct_0(X2)))))), introduced(definition)).
fof(c_0_8, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((~(v2_struct_0(X2))&l1_orders_2(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v23_waybel_0(X3,X1,X2)=>((v1_funct_1(k2_funct_1(X3))&(v1_funct_2(k2_funct_1(X3),u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(k2_funct_1(X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))&k2_relat_1(k2_funct_1(X3))=u1_struct_0(X1))))))), inference(assume_negation,[status(cth)],[t67_waybel_0])).
fof(c_0_9, plain, ![X1, X3, X2]:(epred1_3(X2,X3,X1)=>(~((~(v2_struct_0(X1))&(~(v2_struct_0(X2))&~((v23_waybel_0(X3,X1,X2)<=>(v2_funct_1(X3)&(v5_orders_3(X3,X1,X2)&?[X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))&(X4=k2_funct_1(X3)&v5_orders_3(X4,X2,X1))))))))))&(~((~(v2_struct_0(X1))&~(v2_struct_0(X2))))=>(v23_waybel_0(X3,X1,X2)<=>(v2_struct_0(X1)&v2_struct_0(X2)))))), inference(split_equiv,[status(thm)],[c_0_7])).
fof(c_0_10, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>epred1_3(X2,X3,X1)))), inference(apply_def,[status(thm)],[d38_waybel_0, c_0_7])).
fof(c_0_11, plain, ![X9, X10, X11]:(~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|v1_relat_1(X11)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_12, negated_conjecture, ((~v2_struct_0(esk1_0)&l1_orders_2(esk1_0))&((~v2_struct_0(esk2_0)&l1_orders_2(esk2_0))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))))&(v23_waybel_0(esk3_0,esk1_0,esk2_0)&(~v1_funct_1(k2_funct_1(esk3_0))|(~v1_funct_2(k2_funct_1(esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0)))))|k2_relat_1(k2_funct_1(esk3_0))!=u1_struct_0(esk1_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])).
fof(c_0_13, plain, ![X27, X28, X29, X31]:((((v2_funct_1(X28)|~v23_waybel_0(X28,X27,X29)|v2_struct_0(X29)|v2_struct_0(X27)|~epred1_3(X29,X28,X27))&((v5_orders_3(X28,X27,X29)|~v23_waybel_0(X28,X27,X29)|v2_struct_0(X29)|v2_struct_0(X27)|~epred1_3(X29,X28,X27))&(((v1_funct_1(esk6_3(X27,X28,X29))|~v23_waybel_0(X28,X27,X29)|v2_struct_0(X29)|v2_struct_0(X27)|~epred1_3(X29,X28,X27))&((v1_funct_2(esk6_3(X27,X28,X29),u1_struct_0(X29),u1_struct_0(X27))|~v23_waybel_0(X28,X27,X29)|v2_struct_0(X29)|v2_struct_0(X27)|~epred1_3(X29,X28,X27))&(m1_subset_1(esk6_3(X27,X28,X29),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X29),u1_struct_0(X27))))|~v23_waybel_0(X28,X27,X29)|v2_struct_0(X29)|v2_struct_0(X27)|~epred1_3(X29,X28,X27))))&((esk6_3(X27,X28,X29)=k2_funct_1(X28)|~v23_waybel_0(X28,X27,X29)|v2_struct_0(X29)|v2_struct_0(X27)|~epred1_3(X29,X28,X27))&(v5_orders_3(esk6_3(X27,X28,X29),X29,X27)|~v23_waybel_0(X28,X27,X29)|v2_struct_0(X29)|v2_struct_0(X27)|~epred1_3(X29,X28,X27))))))&(~v2_funct_1(X28)|(~v5_orders_3(X28,X27,X29)|(~v1_funct_1(X31)|(~v1_funct_2(X31,u1_struct_0(X29),u1_struct_0(X27))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X29),u1_struct_0(X27)))))|(X31!=k2_funct_1(X28)|~v5_orders_3(X31,X29,X27))))|v23_waybel_0(X28,X27,X29)|v2_struct_0(X29)|v2_struct_0(X27)|~epred1_3(X29,X28,X27)))&((((v2_struct_0(X27)|~v23_waybel_0(X28,X27,X29)|~v2_struct_0(X27)|~epred1_3(X29,X28,X27))&(v2_struct_0(X29)|~v23_waybel_0(X28,X27,X29)|~v2_struct_0(X27)|~epred1_3(X29,X28,X27)))&(~v2_struct_0(X27)|~v2_struct_0(X29)|v23_waybel_0(X28,X27,X29)|~v2_struct_0(X27)|~epred1_3(X29,X28,X27)))&(((v2_struct_0(X27)|~v23_waybel_0(X28,X27,X29)|~v2_struct_0(X29)|~epred1_3(X29,X28,X27))&(v2_struct_0(X29)|~v23_waybel_0(X28,X27,X29)|~v2_struct_0(X29)|~epred1_3(X29,X28,X27)))&(~v2_struct_0(X27)|~v2_struct_0(X29)|v23_waybel_0(X28,X27,X29)|~v2_struct_0(X29)|~epred1_3(X29,X28,X27))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])])])).
fof(c_0_14, plain, ![X23, X24, X25]:(~l1_orders_2(X23)|(~l1_orders_2(X24)|(~v1_funct_1(X25)|(~v1_funct_2(X25,u1_struct_0(X23),u1_struct_0(X24))|~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X23),u1_struct_0(X24)))))|epred1_3(X24,X25,X23)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
fof(c_0_15, plain, ![X22]:((v1_relat_1(k2_funct_1(X22))|(~v1_relat_1(X22)|(~v1_funct_1(X22)|~v2_funct_1(X22))))&((v1_funct_1(k2_funct_1(X22))|(~v1_relat_1(X22)|(~v1_funct_1(X22)|~v2_funct_1(X22))))&(v2_funct_1(k2_funct_1(X22))|(~v1_relat_1(X22)|(~v1_funct_1(X22)|~v2_funct_1(X22)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc6_funct_1])])])).
cnf(c_0_16, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_18, plain, ![X12, X13, X14, X15, X16]:(((v2_funct_1(X14)|~v23_waybel_0(X14,X12,X13)|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(X12),u1_struct_0(X13))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X12),u1_struct_0(X13))))))|(v2_struct_0(X13)|~l1_orders_2(X13))|(v2_struct_0(X12)|~l1_orders_2(X12)))&((k2_relset_1(u1_struct_0(X12),u1_struct_0(X13),X14)=u1_struct_0(X13)|~v23_waybel_0(X14,X12,X13)|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(X12),u1_struct_0(X13))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X12),u1_struct_0(X13))))))|(v2_struct_0(X13)|~l1_orders_2(X13))|(v2_struct_0(X12)|~l1_orders_2(X12)))&((~r1_orders_2(X12,X15,X16)|r1_orders_2(X13,k3_funct_2(u1_struct_0(X12),u1_struct_0(X13),X14,X15),k3_funct_2(u1_struct_0(X12),u1_struct_0(X13),X14,X16))|~m1_subset_1(X16,u1_struct_0(X12))|~m1_subset_1(X15,u1_struct_0(X12))|~v23_waybel_0(X14,X12,X13)|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(X12),u1_struct_0(X13))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X12),u1_struct_0(X13))))))|(v2_struct_0(X13)|~l1_orders_2(X13))|(v2_struct_0(X12)|~l1_orders_2(X12)))&(~r1_orders_2(X13,k3_funct_2(u1_struct_0(X12),u1_struct_0(X13),X14,X15),k3_funct_2(u1_struct_0(X12),u1_struct_0(X13),X14,X16))|r1_orders_2(X12,X15,X16)|~m1_subset_1(X16,u1_struct_0(X12))|~m1_subset_1(X15,u1_struct_0(X12))|~v23_waybel_0(X14,X12,X13)|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(X12),u1_struct_0(X13))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X12),u1_struct_0(X13))))))|(v2_struct_0(X13)|~l1_orders_2(X13))|(v2_struct_0(X12)|~l1_orders_2(X12))))))&((m1_subset_1(esk4_3(X12,X13,X14),u1_struct_0(X12))|k2_relset_1(u1_struct_0(X12),u1_struct_0(X13),X14)!=u1_struct_0(X13)|~v2_funct_1(X14)|v23_waybel_0(X14,X12,X13)|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(X12),u1_struct_0(X13))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X12),u1_struct_0(X13))))))|(v2_struct_0(X13)|~l1_orders_2(X13))|(v2_struct_0(X12)|~l1_orders_2(X12)))&((m1_subset_1(esk5_3(X12,X13,X14),u1_struct_0(X12))|k2_relset_1(u1_struct_0(X12),u1_struct_0(X13),X14)!=u1_struct_0(X13)|~v2_funct_1(X14)|v23_waybel_0(X14,X12,X13)|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(X12),u1_struct_0(X13))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X12),u1_struct_0(X13))))))|(v2_struct_0(X13)|~l1_orders_2(X13))|(v2_struct_0(X12)|~l1_orders_2(X12)))&((~r1_orders_2(X12,esk4_3(X12,X13,X14),esk5_3(X12,X13,X14))|~r1_orders_2(X13,k3_funct_2(u1_struct_0(X12),u1_struct_0(X13),X14,esk4_3(X12,X13,X14)),k3_funct_2(u1_struct_0(X12),u1_struct_0(X13),X14,esk5_3(X12,X13,X14)))|k2_relset_1(u1_struct_0(X12),u1_struct_0(X13),X14)!=u1_struct_0(X13)|~v2_funct_1(X14)|v23_waybel_0(X14,X12,X13)|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(X12),u1_struct_0(X13))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X12),u1_struct_0(X13))))))|(v2_struct_0(X13)|~l1_orders_2(X13))|(v2_struct_0(X12)|~l1_orders_2(X12)))&(r1_orders_2(X12,esk4_3(X12,X13,X14),esk5_3(X12,X13,X14))|r1_orders_2(X13,k3_funct_2(u1_struct_0(X12),u1_struct_0(X13),X14,esk4_3(X12,X13,X14)),k3_funct_2(u1_struct_0(X12),u1_struct_0(X13),X14,esk5_3(X12,X13,X14)))|k2_relset_1(u1_struct_0(X12),u1_struct_0(X13),X14)!=u1_struct_0(X13)|~v2_funct_1(X14)|v23_waybel_0(X14,X12,X13)|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(X12),u1_struct_0(X13))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X12),u1_struct_0(X13))))))|(v2_struct_0(X13)|~l1_orders_2(X13))|(v2_struct_0(X12)|~l1_orders_2(X12))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t66_waybel_0])])])])])])).
cnf(c_0_19, plain, (esk6_3(X1,X2,X3)=k2_funct_1(X2)|v2_struct_0(X3)|v2_struct_0(X1)|~v23_waybel_0(X2,X1,X3)|~epred1_3(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, plain, (v2_struct_0(X1)|~v23_waybel_0(X2,X3,X1)|~v2_struct_0(X3)|~epred1_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (epred1_3(X2,X3,X1)|~l1_orders_2(X1)|~l1_orders_2(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_24, negated_conjecture, (l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_25, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_26, plain, (v1_funct_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_27, negated_conjecture, (v1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_28, plain, (v2_funct_1(X1)|v2_struct_0(X3)|v2_struct_0(X2)|~v23_waybel_0(X1,X2,X3)|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~l1_orders_2(X3)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_29, negated_conjecture, (v23_waybel_0(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_30, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_31, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_32, plain, (v1_funct_2(esk6_3(X1,X2,X3),u1_struct_0(X3),u1_struct_0(X1))|v2_struct_0(X3)|v2_struct_0(X1)|~v23_waybel_0(X2,X1,X3)|~epred1_3(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_33, plain, (esk6_3(X1,X2,X3)=k2_funct_1(X2)|v2_struct_0(X3)|~epred1_3(X3,X2,X1)|~v23_waybel_0(X2,X1,X3)), inference(csr,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_34, negated_conjecture, (epred1_3(esk2_0,esk3_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_17]), c_0_23]), c_0_24]), c_0_25])])).
cnf(c_0_35, negated_conjecture, (v1_funct_1(k2_funct_1(esk3_0))|~v2_funct_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_23])])).
cnf(c_0_36, negated_conjecture, (v2_funct_1(esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_22]), c_0_29]), c_0_17]), c_0_23]), c_0_24]), c_0_25])]), c_0_30]), c_0_31])).
cnf(c_0_37, plain, (v1_funct_2(esk6_3(X1,X2,X3),u1_struct_0(X3),u1_struct_0(X1))|v2_struct_0(X3)|~epred1_3(X3,X2,X1)|~v23_waybel_0(X2,X1,X3)), inference(csr,[status(thm)],[c_0_32, c_0_20])).
cnf(c_0_38, negated_conjecture, (esk6_3(esk1_0,esk3_0,esk2_0)=k2_funct_1(esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_29])]), c_0_30])).
cnf(c_0_39, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|v2_struct_0(X3)|v2_struct_0(X1)|~v23_waybel_0(X2,X1,X3)|~epred1_3(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_40, plain, ![X26]:(~v1_relat_1(X26)|~v1_funct_1(X26)|(~v2_funct_1(X26)|k2_funct_1(k2_funct_1(X26))=X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t65_funct_1])])).
cnf(c_0_41, negated_conjecture, (~v1_funct_1(k2_funct_1(esk3_0))|~v1_funct_2(k2_funct_1(esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|k2_relat_1(k2_funct_1(esk3_0))!=u1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_42, negated_conjecture, (v1_funct_1(k2_funct_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_36])])).
cnf(c_0_43, plain, (v5_orders_3(esk6_3(X1,X2,X3),X3,X1)|v2_struct_0(X3)|v2_struct_0(X1)|~v23_waybel_0(X2,X1,X3)|~epred1_3(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_44, negated_conjecture, (v1_funct_2(k2_funct_1(esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_34]), c_0_38]), c_0_29])]), c_0_30])).
cnf(c_0_45, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|v2_struct_0(X3)|~epred1_3(X3,X2,X1)|~v23_waybel_0(X2,X1,X3)), inference(csr,[status(thm)],[c_0_39, c_0_20])).
cnf(c_0_46, plain, (k2_funct_1(k2_funct_1(X1))=X1|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_47, plain, (v5_orders_3(X1,X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|~v23_waybel_0(X1,X2,X3)|~epred1_3(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_48, plain, (v2_funct_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_49, negated_conjecture, (k2_relat_1(k2_funct_1(esk3_0))!=u1_struct_0(esk1_0)|~m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(k2_funct_1(esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_42])])).
fof(c_0_50, plain, ![X19, X20, X21]:(~m1_subset_1(X21,k1_zfmisc_1(k2_zfmisc_1(X19,X20)))|k2_relset_1(X19,X20,X21)=k2_relat_1(X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
cnf(c_0_51, plain, (v23_waybel_0(X1,X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|~v2_funct_1(X1)|~v5_orders_3(X1,X2,X3)|~v1_funct_1(X4)|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|X4!=k2_funct_1(X1)|~v5_orders_3(X4,X3,X2)|~epred1_3(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_52, plain, (v5_orders_3(esk6_3(X1,X2,X3),X3,X1)|v2_struct_0(X3)|~epred1_3(X3,X2,X1)|~v23_waybel_0(X2,X1,X3)), inference(csr,[status(thm)],[c_0_43, c_0_20])).
cnf(c_0_53, negated_conjecture, (epred1_3(esk1_0,k2_funct_1(esk3_0),esk2_0)|~m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_44]), c_0_42]), c_0_25]), c_0_24])])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_34]), c_0_38]), c_0_29])]), c_0_30])).
cnf(c_0_55, negated_conjecture, (k2_funct_1(k2_funct_1(esk3_0))=esk3_0|~v2_funct_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_27]), c_0_23])])).
cnf(c_0_56, plain, (v5_orders_3(X1,X2,X3)|v2_struct_0(X3)|~epred1_3(X3,X1,X2)|~v23_waybel_0(X1,X2,X3)), inference(csr,[status(thm)],[c_0_47, c_0_20])).
cnf(c_0_57, negated_conjecture, (v2_funct_1(k2_funct_1(esk3_0))|~v2_funct_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_27]), c_0_23])])).
cnf(c_0_58, negated_conjecture, (k2_relat_1(k2_funct_1(esk3_0))!=u1_struct_0(esk1_0)|~m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_44])])).
cnf(c_0_59, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_60, plain, (k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=u1_struct_0(X2)|v2_struct_0(X2)|v2_struct_0(X1)|~v23_waybel_0(X3,X1,X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_orders_2(X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_61, plain, (v23_waybel_0(X1,X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|~epred1_3(X3,X1,X2)|~v5_orders_3(k2_funct_1(X1),X3,X2)|~v5_orders_3(X1,X2,X3)|~v2_funct_1(X1)|~m1_subset_1(k2_funct_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(k2_funct_1(X1),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(k2_funct_1(X1))), inference(er,[status(thm)],[c_0_51])).
cnf(c_0_62, negated_conjecture, (v5_orders_3(k2_funct_1(esk3_0),esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_34]), c_0_38]), c_0_29])]), c_0_30])).
cnf(c_0_63, negated_conjecture, (epred1_3(esk1_0,k2_funct_1(esk3_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_54])])).
cnf(c_0_64, negated_conjecture, (k2_funct_1(k2_funct_1(esk3_0))=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_36])])).
cnf(c_0_65, negated_conjecture, (v5_orders_3(esk3_0,esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_34]), c_0_29])]), c_0_30])).
cnf(c_0_66, negated_conjecture, (v2_funct_1(k2_funct_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_36])])).
cnf(c_0_67, negated_conjecture, (k2_relat_1(k2_funct_1(esk3_0))!=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_54])])).
cnf(c_0_68, negated_conjecture, (k2_relat_1(k2_funct_1(esk3_0))=k2_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_funct_1(esk3_0))), inference(spm,[status(thm)],[c_0_59, c_0_54])).
cnf(c_0_69, negated_conjecture, (k2_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_funct_1(esk3_0))=u1_struct_0(esk1_0)|~v23_waybel_0(k2_funct_1(esk3_0),esk2_0,esk1_0)|~m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_44]), c_0_42]), c_0_25]), c_0_24])]), c_0_31]), c_0_30])).
cnf(c_0_70, plain, (v23_waybel_0(k2_funct_1(esk3_0),esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63]), c_0_64]), c_0_65]), c_0_66]), c_0_64]), c_0_17]), c_0_64]), c_0_22]), c_0_64]), c_0_23])]), c_0_30]), c_0_31])).
cnf(c_0_71, negated_conjecture, (k2_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_funct_1(esk3_0))!=u1_struct_0(esk1_0)), inference(rw,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_72, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_54])]), c_0_70])]), c_0_71]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 73
# Proof object clause steps            : 54
# Proof object formula steps           : 19
# Proof object conjectures             : 35
# Proof object clause conjectures      : 32
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 86
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 38
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 36
# Processed clauses                    : 133
# ...of these trivial                  : 2
# ...subsumed                          : 6
# ...remaining for further processing  : 124
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 13
# Generated clauses                    : 113
# ...of the previous two non-trivial   : 68
# Contextual simplify-reflections      : 51
# Paramodulations                      : 112
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 75
#    Positive orientable unit clauses  : 24
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 48
# Current number of unprocessed clauses: 6
# ...number of literals in the above   : 24
# Current number of archived formulas  : 0
# Current number of archived clauses   : 48
# Clause-clause subsumption calls (NU) : 2176
# Rec. Clause-clause subsumption calls : 448
# Non-unit clause-clause subsumptions  : 52
# Unit Clause-clause subsumption calls : 3
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7885

# -------------------------------------------------
# User time                : 0.014 s
# System time              : 0.000 s
# Total time               : 0.014 s
# Maximum resident set size: 29020 pages
