# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>(v7_struct_0(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>X2=X3)))), file('data/MPTP/proofs/33d821085d7943f48a4fd329a3c52590.E_input', d10_struct_0)).
fof(dt_k4_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(k4_yellow_0(X1),u1_struct_0(X1))), file('data/MPTP/proofs/33d821085d7943f48a4fd329a3c52590.E_input', dt_k4_yellow_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/33d821085d7943f48a4fd329a3c52590.E_input', dt_l1_orders_2)).
fof(dt_k3_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(k3_yellow_0(X1),u1_struct_0(X1))), file('data/MPTP/proofs/33d821085d7943f48a4fd329a3c52590.E_input', dt_k3_yellow_0)).
fof(t5_waybel_7, conjecture, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v3_yellow_0(X1)&l1_orders_2(X1))))=>(v7_struct_0(X1)<=>k4_yellow_0(X1)=k3_yellow_0(X1))), file('data/MPTP/proofs/33d821085d7943f48a4fd329a3c52590.E_input', t5_waybel_7)).
fof(t44_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,k3_yellow_0(X1),X2))), file('data/MPTP/proofs/33d821085d7943f48a4fd329a3c52590.E_input', t44_yellow_0)).
fof(d10_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_yellow_0(X1,X2)=>(X3=k2_yellow_0(X1,X2)<=>(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))))))), file('data/MPTP/proofs/33d821085d7943f48a4fd329a3c52590.E_input', d10_yellow_0)).
fof(t25_orders_2, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X3,X2))=>X2=X3)))), file('data/MPTP/proofs/33d821085d7943f48a4fd329a3c52590.E_input', t25_orders_2)).
fof(d12_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>k4_yellow_0(X1)=k2_yellow_0(X1,k1_xboole_0)), file('data/MPTP/proofs/33d821085d7943f48a4fd329a3c52590.E_input', d12_yellow_0)).
fof(t6_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_lattice3(X1,k1_xboole_0,X2)&r1_lattice3(X1,k1_xboole_0,X2)))), file('data/MPTP/proofs/33d821085d7943f48a4fd329a3c52590.E_input', t6_yellow_0)).
fof(t43_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v2_yellow_0(X1)&l1_orders_2(X1))))=>(r2_yellow_0(X1,k1_xboole_0)&r1_yellow_0(X1,u1_struct_0(X1)))), file('data/MPTP/proofs/33d821085d7943f48a4fd329a3c52590.E_input', t43_yellow_0)).
fof(cc4_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>(v3_yellow_0(X1)=>(v1_yellow_0(X1)&v2_yellow_0(X1)))), file('data/MPTP/proofs/33d821085d7943f48a4fd329a3c52590.E_input', cc4_yellow_0)).
fof(c_0_12, plain, ![X5, X6, X7]:((~v7_struct_0(X5)|(~m1_subset_1(X6,u1_struct_0(X5))|(~m1_subset_1(X7,u1_struct_0(X5))|X6=X7))|~l1_struct_0(X5))&((m1_subset_1(esk1_1(X5),u1_struct_0(X5))|v7_struct_0(X5)|~l1_struct_0(X5))&((m1_subset_1(esk2_1(X5),u1_struct_0(X5))|v7_struct_0(X5)|~l1_struct_0(X5))&(esk1_1(X5)!=esk2_1(X5)|v7_struct_0(X5)|~l1_struct_0(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_struct_0])])])])])).
fof(c_0_13, plain, ![X17]:(~l1_orders_2(X17)|m1_subset_1(k4_yellow_0(X17),u1_struct_0(X17))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_yellow_0])])).
fof(c_0_14, plain, ![X11]:(~l1_orders_2(X11)|l1_struct_0(X11)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_15, plain, (X2=X3|~v7_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, plain, (m1_subset_1(k4_yellow_0(X1),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, plain, ![X10]:(~l1_orders_2(X10)|m1_subset_1(k3_yellow_0(X10),u1_struct_0(X10))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_yellow_0])])).
fof(c_0_19, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v3_yellow_0(X1)&l1_orders_2(X1))))=>(v7_struct_0(X1)<=>k4_yellow_0(X1)=k3_yellow_0(X1)))), inference(assume_negation,[status(cth)],[t5_waybel_7])).
cnf(c_0_20, plain, (X1=k4_yellow_0(X2)|~l1_orders_2(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~v7_struct_0(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])).
cnf(c_0_21, plain, (m1_subset_1(k3_yellow_0(X1),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_22, negated_conjecture, ((~v2_struct_0(esk4_0)&(v5_orders_2(esk4_0)&(v3_yellow_0(esk4_0)&l1_orders_2(esk4_0))))&((~v7_struct_0(esk4_0)|k4_yellow_0(esk4_0)!=k3_yellow_0(esk4_0))&(v7_struct_0(esk4_0)|k4_yellow_0(esk4_0)=k3_yellow_0(esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_19])])])])).
fof(c_0_23, plain, ![X25, X26]:(v2_struct_0(X25)|(~v5_orders_2(X25)|(~v1_yellow_0(X25)|~l1_orders_2(X25)))|(~m1_subset_1(X26,u1_struct_0(X25))|r1_orders_2(X25,k3_yellow_0(X25),X26))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t44_yellow_0])])])])).
cnf(c_0_24, plain, (k3_yellow_0(X1)=k4_yellow_0(X1)|~l1_orders_2(X1)|~v7_struct_0(X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_25, negated_conjecture, (v7_struct_0(esk4_0)|k4_yellow_0(esk4_0)=k3_yellow_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, negated_conjecture, (l1_orders_2(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_27, plain, ![X12, X13, X14, X15]:(((r1_lattice3(X12,X13,X14)|X14!=k2_yellow_0(X12,X13)|~r2_yellow_0(X12,X13)|~m1_subset_1(X14,u1_struct_0(X12))|~l1_orders_2(X12))&(~m1_subset_1(X15,u1_struct_0(X12))|(~r1_lattice3(X12,X13,X15)|r1_orders_2(X12,X15,X14))|X14!=k2_yellow_0(X12,X13)|~r2_yellow_0(X12,X13)|~m1_subset_1(X14,u1_struct_0(X12))|~l1_orders_2(X12)))&((m1_subset_1(esk3_3(X12,X13,X14),u1_struct_0(X12))|~r1_lattice3(X12,X13,X14)|X14=k2_yellow_0(X12,X13)|~r2_yellow_0(X12,X13)|~m1_subset_1(X14,u1_struct_0(X12))|~l1_orders_2(X12))&((r1_lattice3(X12,X13,esk3_3(X12,X13,X14))|~r1_lattice3(X12,X13,X14)|X14=k2_yellow_0(X12,X13)|~r2_yellow_0(X12,X13)|~m1_subset_1(X14,u1_struct_0(X12))|~l1_orders_2(X12))&(~r1_orders_2(X12,esk3_3(X12,X13,X14),X14)|~r1_lattice3(X12,X13,X14)|X14=k2_yellow_0(X12,X13)|~r2_yellow_0(X12,X13)|~m1_subset_1(X14,u1_struct_0(X12))|~l1_orders_2(X12))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_yellow_0])])])])])).
fof(c_0_28, plain, ![X22, X23, X24]:(~v5_orders_2(X22)|~l1_orders_2(X22)|(~m1_subset_1(X23,u1_struct_0(X22))|(~m1_subset_1(X24,u1_struct_0(X22))|(~r1_orders_2(X22,X23,X24)|~r1_orders_2(X22,X24,X23)|X23=X24)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t25_orders_2])])])).
cnf(c_0_29, plain, (v2_struct_0(X1)|r1_orders_2(X1,k3_yellow_0(X1),X2)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30, negated_conjecture, (k3_yellow_0(esk4_0)=k4_yellow_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26])])).
cnf(c_0_31, negated_conjecture, (v5_orders_2(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33, plain, (r1_orders_2(X2,X1,X4)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|X4!=k2_yellow_0(X2,X3)|~r2_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_34, plain, ![X19]:(~l1_orders_2(X19)|k4_yellow_0(X19)=k2_yellow_0(X19,k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_yellow_0])])).
fof(c_0_35, plain, ![X20, X21]:((r2_lattice3(X20,k1_xboole_0,X21)|~m1_subset_1(X21,u1_struct_0(X20))|~l1_orders_2(X20))&(r1_lattice3(X20,k1_xboole_0,X21)|~m1_subset_1(X21,u1_struct_0(X20))|~l1_orders_2(X20))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_yellow_0])])])])).
cnf(c_0_36, plain, (X2=X3|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, negated_conjecture, (r1_orders_2(esk4_0,k4_yellow_0(esk4_0),X1)|~v1_yellow_0(esk4_0)|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_26])]), c_0_32])).
cnf(c_0_38, negated_conjecture, (m1_subset_1(k4_yellow_0(esk4_0),u1_struct_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_30]), c_0_26])])).
cnf(c_0_39, plain, (r1_orders_2(X1,X2,k2_yellow_0(X1,X3))|~r1_lattice3(X1,X3,X2)|~r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(er,[status(thm)],[c_0_33])).
cnf(c_0_40, plain, (k4_yellow_0(X1)=k2_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_41, plain, (r1_lattice3(X1,k1_xboole_0,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_42, negated_conjecture, (X1=k4_yellow_0(esk4_0)|~v1_yellow_0(esk4_0)|~r1_orders_2(esk4_0,X1,k4_yellow_0(esk4_0))|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_31]), c_0_26])]), c_0_38])])).
cnf(c_0_43, plain, (r1_orders_2(X1,X2,k4_yellow_0(X1))|~r2_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_16]), c_0_41])).
cnf(c_0_44, negated_conjecture, (~v7_struct_0(esk4_0)|k4_yellow_0(esk4_0)!=k3_yellow_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_45, negated_conjecture, (X1=k4_yellow_0(esk4_0)|~v1_yellow_0(esk4_0)|~r2_yellow_0(esk4_0,k1_xboole_0)|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_26])])).
cnf(c_0_46, plain, (m1_subset_1(esk2_1(X1),u1_struct_0(X1))|v7_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_47, negated_conjecture, (~v7_struct_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_30])])).
fof(c_0_48, plain, ![X27]:((r2_yellow_0(X27,k1_xboole_0)|(v2_struct_0(X27)|(~v5_orders_2(X27)|(~v2_yellow_0(X27)|~l1_orders_2(X27)))))&(r1_yellow_0(X27,u1_struct_0(X27))|(v2_struct_0(X27)|(~v5_orders_2(X27)|(~v2_yellow_0(X27)|~l1_orders_2(X27)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t43_yellow_0])])])])).
cnf(c_0_49, negated_conjecture, (esk2_1(esk4_0)=k4_yellow_0(esk4_0)|~v1_yellow_0(esk4_0)|~r2_yellow_0(esk4_0,k1_xboole_0)|~l1_struct_0(esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])).
cnf(c_0_50, plain, (r2_yellow_0(X1,k1_xboole_0)|v2_struct_0(X1)|~v5_orders_2(X1)|~v2_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_51, plain, ![X28]:((v1_yellow_0(X28)|~v3_yellow_0(X28)|~l1_orders_2(X28))&(v2_yellow_0(X28)|~v3_yellow_0(X28)|~l1_orders_2(X28))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_yellow_0])])])).
cnf(c_0_52, negated_conjecture, (esk2_1(esk4_0)=k4_yellow_0(esk4_0)|~v2_yellow_0(esk4_0)|~v1_yellow_0(esk4_0)|~l1_struct_0(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_31]), c_0_26])]), c_0_32])).
cnf(c_0_53, plain, (v2_yellow_0(X1)|~v3_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_54, negated_conjecture, (v3_yellow_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_55, negated_conjecture, (esk2_1(esk4_0)=k4_yellow_0(esk4_0)|~v1_yellow_0(esk4_0)|~l1_struct_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54]), c_0_26])])).
cnf(c_0_56, plain, (v1_yellow_0(X1)|~v3_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_57, plain, (v7_struct_0(X1)|esk1_1(X1)!=esk2_1(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_58, negated_conjecture, (esk2_1(esk4_0)=k4_yellow_0(esk4_0)|~l1_struct_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_54]), c_0_26])])).
cnf(c_0_59, plain, (m1_subset_1(esk1_1(X1),u1_struct_0(X1))|v7_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_60, negated_conjecture, (esk1_1(esk4_0)!=k4_yellow_0(esk4_0)|~l1_struct_0(esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_47])).
cnf(c_0_61, negated_conjecture, (~v1_yellow_0(esk4_0)|~r2_yellow_0(esk4_0,k1_xboole_0)|~l1_struct_0(esk4_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_59]), c_0_47]), c_0_60])).
cnf(c_0_62, negated_conjecture, (~v2_yellow_0(esk4_0)|~v1_yellow_0(esk4_0)|~l1_struct_0(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_50]), c_0_31]), c_0_26])]), c_0_32])).
cnf(c_0_63, negated_conjecture, (~v1_yellow_0(esk4_0)|~l1_struct_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_53]), c_0_54]), c_0_26])])).
cnf(c_0_64, negated_conjecture, (~l1_struct_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_56]), c_0_54]), c_0_26])])).
cnf(c_0_65, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_17]), c_0_26])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 66
# Proof object clause steps            : 41
# Proof object formula steps           : 25
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 47
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 54
# ...of these trivial                  : 0
# ...subsumed                          : 4
# ...remaining for further processing  : 50
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 2
# Generated clauses                    : 52
# ...of the previous two non-trivial   : 38
# Contextual simplify-reflections      : 7
# Paramodulations                      : 50
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 41
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 33
# Current number of unprocessed clauses: 11
# ...number of literals in the above   : 82
# Current number of archived formulas  : 0
# Current number of archived clauses   : 7
# Clause-clause subsumption calls (NU) : 386
# Rec. Clause-clause subsumption calls : 67
# Non-unit clause-clause subsumptions  : 15
# Unit Clause-clause subsumption calls : 3
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3047

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28652 pages
