# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t17_yellow19, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_waybel_9(X1,k3_yellow19(X1,k2_struct_0(X1),X2),X3)<=>r1_waybel_7(X1,X2,X3))))), file('data/MPTP/proofs/3403973aef7a4814bb27495a170b0fc9.E_input', t17_yellow19)).
fof(dt_k3_yellow19, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(~(v1_xboole_0(X3))&(v2_waybel_0(X3,k3_yellow_1(X2))&(v13_waybel_0(X3,k3_yellow_1(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))))))=>(~(v2_struct_0(k3_yellow19(X1,X2,X3)))&(v6_waybel_0(k3_yellow19(X1,X2,X3),X1)&l1_waybel_0(k3_yellow19(X1,X2,X3),X1)))), file('data/MPTP/proofs/3403973aef7a4814bb27495a170b0fc9.E_input', dt_k3_yellow19)).
fof(d2_yellow_1, axiom, ![X1]:k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1)), file('data/MPTP/proofs/3403973aef7a4814bb27495a170b0fc9.E_input', d2_yellow_1)).
fof(fc4_yellow19, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(~(v1_xboole_0(X3))&(v2_waybel_0(X3,k3_yellow_1(X2))&(v13_waybel_0(X3,k3_yellow_1(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))))))=>(~(v2_struct_0(k3_yellow19(X1,X2,X3)))&(v3_orders_2(k3_yellow19(X1,X2,X3))&(v4_orders_2(k3_yellow19(X1,X2,X3))&v6_waybel_0(k3_yellow19(X1,X2,X3),X1))))), file('data/MPTP/proofs/3403973aef7a4814bb27495a170b0fc9.E_input', fc4_yellow19)).
fof(fc5_yellow19, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(~(v1_xboole_0(X3))&(v1_subset_1(X3,u1_struct_0(k3_yellow_1(X2)))&(v2_waybel_0(X3,k3_yellow_1(X2))&(v13_waybel_0(X3,k3_yellow_1(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))))))))=>(~(v2_struct_0(k3_yellow19(X1,X2,X3)))&(v6_waybel_0(k3_yellow19(X1,X2,X3),X1)&v7_waybel_0(k3_yellow19(X1,X2,X3))))), file('data/MPTP/proofs/3403973aef7a4814bb27495a170b0fc9.E_input', fc5_yellow19)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/3403973aef7a4814bb27495a170b0fc9.E_input', dt_l1_pre_topc)).
fof(t15_yellow19, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2)))), file('data/MPTP/proofs/3403973aef7a4814bb27495a170b0fc9.E_input', t15_yellow19)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/3403973aef7a4814bb27495a170b0fc9.E_input', d5_xboole_0)).
fof(t12_yellow19, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_waybel_9(X1,X2,X3)<=>r1_waybel_7(X1,k2_yellow19(X1,X2),X3))))), file('data/MPTP/proofs/3403973aef7a4814bb27495a170b0fc9.E_input', t12_yellow19)).
fof(dt_k2_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/3403973aef7a4814bb27495a170b0fc9.E_input', dt_k2_struct_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/3403973aef7a4814bb27495a170b0fc9.E_input', d1_xboole_0)).
fof(redefinition_k7_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k7_subset_1(X1,X2,X3)=k4_xboole_0(X2,X3)), file('data/MPTP/proofs/3403973aef7a4814bb27495a170b0fc9.E_input', redefinition_k7_subset_1)).
fof(t22_pre_topc, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=X2)), file('data/MPTP/proofs/3403973aef7a4814bb27495a170b0fc9.E_input', t22_pre_topc)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/3403973aef7a4814bb27495a170b0fc9.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/3403973aef7a4814bb27495a170b0fc9.E_input', d4_subset_1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/3403973aef7a4814bb27495a170b0fc9.E_input', fc2_struct_0)).
fof(c_0_16, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_waybel_9(X1,k3_yellow19(X1,k2_struct_0(X1),X2),X3)<=>r1_waybel_7(X1,X2,X3)))))), inference(assume_negation,[status(cth)],[t17_yellow19])).
fof(c_0_17, plain, ![X5, X6, X7]:((~v2_struct_0(k3_yellow19(X5,X6,X7))|(v2_struct_0(X5)|~l1_struct_0(X5)|(v1_xboole_0(X6)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))|(v1_xboole_0(X7)|(~v2_waybel_0(X7,k3_yellow_1(X6))|(~v13_waybel_0(X7,k3_yellow_1(X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X6))))))))))&((v6_waybel_0(k3_yellow19(X5,X6,X7),X5)|(v2_struct_0(X5)|~l1_struct_0(X5)|(v1_xboole_0(X6)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))|(v1_xboole_0(X7)|(~v2_waybel_0(X7,k3_yellow_1(X6))|(~v13_waybel_0(X7,k3_yellow_1(X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X6))))))))))&(l1_waybel_0(k3_yellow19(X5,X6,X7),X5)|(v2_struct_0(X5)|~l1_struct_0(X5)|(v1_xboole_0(X6)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))|(v1_xboole_0(X7)|(~v2_waybel_0(X7,k3_yellow_1(X6))|(~v13_waybel_0(X7,k3_yellow_1(X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X6)))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_yellow19])])])])).
fof(c_0_18, plain, ![X8]:k3_yellow_1(X8)=k3_lattice3(k1_lattice3(X8)), inference(variable_rename,[status(thm)],[d2_yellow_1])).
fof(c_0_19, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v1_xboole_0(esk2_0)&(v1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))&(v2_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))&(v13_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))))))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&((~r3_waybel_9(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk3_0)|~r1_waybel_7(esk1_0,esk2_0,esk3_0))&(r3_waybel_9(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk3_0)|r1_waybel_7(esk1_0,esk2_0,esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_16])])])])).
fof(c_0_20, plain, ![X12, X13, X14]:((~v2_struct_0(k3_yellow19(X12,X13,X14))|(v2_struct_0(X12)|~l1_struct_0(X12)|(v1_xboole_0(X13)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|(v1_xboole_0(X14)|(~v2_waybel_0(X14,k3_yellow_1(X13))|(~v13_waybel_0(X14,k3_yellow_1(X13))|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X13))))))))))&((v3_orders_2(k3_yellow19(X12,X13,X14))|(v2_struct_0(X12)|~l1_struct_0(X12)|(v1_xboole_0(X13)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|(v1_xboole_0(X14)|(~v2_waybel_0(X14,k3_yellow_1(X13))|(~v13_waybel_0(X14,k3_yellow_1(X13))|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X13))))))))))&((v4_orders_2(k3_yellow19(X12,X13,X14))|(v2_struct_0(X12)|~l1_struct_0(X12)|(v1_xboole_0(X13)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|(v1_xboole_0(X14)|(~v2_waybel_0(X14,k3_yellow_1(X13))|(~v13_waybel_0(X14,k3_yellow_1(X13))|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X13))))))))))&(v6_waybel_0(k3_yellow19(X12,X13,X14),X12)|(v2_struct_0(X12)|~l1_struct_0(X12)|(v1_xboole_0(X13)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|(v1_xboole_0(X14)|(~v2_waybel_0(X14,k3_yellow_1(X13))|(~v13_waybel_0(X14,k3_yellow_1(X13))|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X13))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc4_yellow19])])])])).
fof(c_0_21, plain, ![X19, X20, X21]:((~v2_struct_0(k3_yellow19(X19,X20,X21))|(v2_struct_0(X19)|~l1_struct_0(X19)|(v1_xboole_0(X20)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|(v1_xboole_0(X21)|(~v1_subset_1(X21,u1_struct_0(k3_yellow_1(X20)))|(~v2_waybel_0(X21,k3_yellow_1(X20))|(~v13_waybel_0(X21,k3_yellow_1(X20))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X20)))))))))))&((v6_waybel_0(k3_yellow19(X19,X20,X21),X19)|(v2_struct_0(X19)|~l1_struct_0(X19)|(v1_xboole_0(X20)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|(v1_xboole_0(X21)|(~v1_subset_1(X21,u1_struct_0(k3_yellow_1(X20)))|(~v2_waybel_0(X21,k3_yellow_1(X20))|(~v13_waybel_0(X21,k3_yellow_1(X20))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X20)))))))))))&(v7_waybel_0(k3_yellow19(X19,X20,X21))|(v2_struct_0(X19)|~l1_struct_0(X19)|(v1_xboole_0(X20)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|(v1_xboole_0(X21)|(~v1_subset_1(X21,u1_struct_0(k3_yellow_1(X20)))|(~v2_waybel_0(X21,k3_yellow_1(X20))|(~v13_waybel_0(X21,k3_yellow_1(X20))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X20))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc5_yellow19])])])])).
cnf(c_0_22, plain, (l1_waybel_0(k3_yellow19(X1,X2,X3),X1)|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, negated_conjecture, (v13_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, negated_conjecture, (v2_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_27, plain, ![X16]:(~l1_pre_topc(X16)|l1_struct_0(X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_28, plain, ![X17, X18]:(v2_struct_0(X17)|~l1_struct_0(X17)|(v1_xboole_0(X18)|(~v1_subset_1(X18,u1_struct_0(k3_yellow_1(k2_struct_0(X17))))|(~v2_waybel_0(X18,k3_yellow_1(k2_struct_0(X17)))|(~v13_waybel_0(X18,k3_yellow_1(k2_struct_0(X17)))|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X17))))))))|X18=k2_yellow19(X17,k3_yellow19(X17,k2_struct_0(X17),X18)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t15_yellow19])])])])).
cnf(c_0_29, plain, (v4_orders_2(k3_yellow19(X1,X2,X3))|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30, plain, (v7_waybel_0(k3_yellow19(X1,X2,X3))|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_subset_1(X3,u1_struct_0(k3_yellow_1(X2)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_32, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~v2_struct_0(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_33, plain, ![X22, X23, X24, X25, X26, X27, X28, X29]:((((r2_hidden(X25,X22)|~r2_hidden(X25,X24)|X24!=k4_xboole_0(X22,X23))&(~r2_hidden(X25,X23)|~r2_hidden(X25,X24)|X24!=k4_xboole_0(X22,X23)))&(~r2_hidden(X26,X22)|r2_hidden(X26,X23)|r2_hidden(X26,X24)|X24!=k4_xboole_0(X22,X23)))&((~r2_hidden(esk4_3(X27,X28,X29),X29)|(~r2_hidden(esk4_3(X27,X28,X29),X27)|r2_hidden(esk4_3(X27,X28,X29),X28))|X29=k4_xboole_0(X27,X28))&((r2_hidden(esk4_3(X27,X28,X29),X27)|r2_hidden(esk4_3(X27,X28,X29),X29)|X29=k4_xboole_0(X27,X28))&(~r2_hidden(esk4_3(X27,X28,X29),X28)|r2_hidden(esk4_3(X27,X28,X29),X29)|X29=k4_xboole_0(X27,X28))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_34, plain, ![X31, X32, X33]:((~r3_waybel_9(X31,X32,X33)|r1_waybel_7(X31,k2_yellow19(X31,X32),X33)|~m1_subset_1(X33,u1_struct_0(X31))|(v2_struct_0(X32)|(~v4_orders_2(X32)|(~v7_waybel_0(X32)|~l1_waybel_0(X32,X31))))|(v2_struct_0(X31)|(~v2_pre_topc(X31)|~l1_pre_topc(X31))))&(~r1_waybel_7(X31,k2_yellow19(X31,X32),X33)|r3_waybel_9(X31,X32,X33)|~m1_subset_1(X33,u1_struct_0(X31))|(v2_struct_0(X32)|(~v4_orders_2(X32)|(~v7_waybel_0(X32)|~l1_waybel_0(X32,X31))))|(v2_struct_0(X31)|(~v2_pre_topc(X31)|~l1_pre_topc(X31))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t12_yellow19])])])])])).
cnf(c_0_35, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|l1_waybel_0(k3_yellow19(X1,X2,X3),X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_23]), c_0_23]), c_0_23])).
cnf(c_0_36, negated_conjecture, (v13_waybel_0(esk2_0,k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))), inference(rw,[status(thm)],[c_0_24, c_0_23])).
cnf(c_0_37, negated_conjecture, (v2_waybel_0(esk2_0,k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))), inference(rw,[status(thm)],[c_0_25, c_0_23])).
cnf(c_0_38, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))))), inference(rw,[status(thm)],[c_0_26, c_0_23])).
cnf(c_0_39, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_40, plain, ![X15]:(~l1_struct_0(X15)|m1_subset_1(k2_struct_0(X15),k1_zfmisc_1(u1_struct_0(X15)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_struct_0])])).
cnf(c_0_41, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_42, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_43, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2))|~l1_struct_0(X1)|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_44, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|v4_orders_2(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_23]), c_0_23]), c_0_23])).
cnf(c_0_45, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|v7_waybel_0(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v1_subset_1(X3,u1_struct_0(k3_lattice3(k1_lattice3(X2))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_23]), c_0_23]), c_0_23]), c_0_23])).
cnf(c_0_46, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(esk1_0)))))), inference(rw,[status(thm)],[c_0_31, c_0_23])).
cnf(c_0_47, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_struct_0(k3_yellow19(X1,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_23]), c_0_23]), c_0_23])).
cnf(c_0_48, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_49, plain, ![X34, X35, X36]:((~v1_xboole_0(X34)|~r2_hidden(X35,X34))&(r2_hidden(esk5_1(X36),X36)|v1_xboole_0(X36))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_50, plain, (r3_waybel_9(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_waybel_7(X1,k2_yellow19(X1,X2),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_51, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_52, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_53, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_54, negated_conjecture, (l1_waybel_0(k3_yellow19(X1,k2_struct_0(esk1_0),esk2_0),X1)|v1_xboole_0(k2_struct_0(esk1_0))|v2_struct_0(X1)|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_55, plain, (m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_56, negated_conjecture, (l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_57, plain, (X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2))|v2_struct_0(X1)|v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~v1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(X1)))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(X1))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_23]), c_0_23]), c_0_23]), c_0_23])).
cnf(c_0_58, negated_conjecture, (v4_orders_2(k3_yellow19(X1,k2_struct_0(esk1_0),esk2_0))|v1_xboole_0(k2_struct_0(esk1_0))|v2_struct_0(X1)|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_36]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_59, negated_conjecture, (v7_waybel_0(k3_yellow19(X1,k2_struct_0(esk1_0),esk2_0))|v1_xboole_0(k2_struct_0(esk1_0))|v2_struct_0(X1)|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_36]), c_0_46]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_60, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))|v2_struct_0(X1)|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)|~v2_struct_0(k3_yellow19(X1,k2_struct_0(esk1_0),esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_36]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_61, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_48])).
cnf(c_0_62, plain, (r2_hidden(esk5_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_63, plain, ![X38, X39, X40]:(~m1_subset_1(X39,k1_zfmisc_1(X38))|k7_subset_1(X38,X39,X40)=k4_xboole_0(X39,X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_subset_1])])).
fof(c_0_64, plain, ![X41, X42]:(~l1_struct_0(X41)|(~m1_subset_1(X42,k1_zfmisc_1(u1_struct_0(X41)))|k7_subset_1(u1_struct_0(X41),k2_struct_0(X41),k7_subset_1(u1_struct_0(X41),k2_struct_0(X41),X42))=X42)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t22_pre_topc])])])).
cnf(c_0_65, negated_conjecture, (r3_waybel_9(esk1_0,X1,esk3_0)|v2_struct_0(X1)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~r1_waybel_7(esk1_0,k2_yellow19(esk1_0,X1),esk3_0)|~l1_waybel_0(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_42]), c_0_52])]), c_0_53])).
cnf(c_0_66, negated_conjecture, (l1_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk1_0)|v1_xboole_0(k2_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56])]), c_0_53])).
cnf(c_0_67, negated_conjecture, (k2_yellow19(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0))=esk2_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_36]), c_0_46]), c_0_37]), c_0_38]), c_0_56])]), c_0_39]), c_0_53])).
cnf(c_0_68, negated_conjecture, (r3_waybel_9(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk3_0)|r1_waybel_7(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_69, negated_conjecture, (v4_orders_2(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0))|v1_xboole_0(k2_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_55]), c_0_56])]), c_0_53])).
cnf(c_0_70, negated_conjecture, (v7_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0))|v1_xboole_0(k2_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_55]), c_0_56])]), c_0_53])).
cnf(c_0_71, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))|~v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_55]), c_0_56])]), c_0_53])).
cnf(c_0_72, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_73, plain, (r2_hidden(esk5_1(k4_xboole_0(X1,X2)),X1)|v1_xboole_0(k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_74, plain, (k7_subset_1(X2,X1,X3)=k4_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_75, plain, (k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=X2|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_76, plain, (r1_waybel_7(X1,k2_yellow19(X1,X2),X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_77, negated_conjecture, (~r3_waybel_9(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk3_0)|~r1_waybel_7(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_78, negated_conjecture, (r3_waybel_9(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk3_0)|v1_xboole_0(k2_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67]), c_0_68]), c_0_69]), c_0_70]), c_0_71])).
cnf(c_0_79, plain, (v1_xboole_0(k4_xboole_0(X1,X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_80, plain, (k4_xboole_0(k2_struct_0(X1),k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=X2|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_55])).
cnf(c_0_81, negated_conjecture, (r1_waybel_7(esk1_0,k2_yellow19(esk1_0,X1),esk3_0)|v2_struct_0(X1)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~r3_waybel_9(esk1_0,X1,esk3_0)|~l1_waybel_0(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_51]), c_0_42]), c_0_52])]), c_0_53])).
cnf(c_0_82, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))|~r1_waybel_7(esk1_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
fof(c_0_83, plain, ![X43]:m1_subset_1(k2_subset_1(X43),k1_zfmisc_1(X43)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_84, plain, ![X44]:k2_subset_1(X44)=X44, inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_85, plain, ![X45]:(v2_struct_0(X45)|~l1_struct_0(X45)|~v1_xboole_0(u1_struct_0(X45))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_86, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_xboole_0(k2_struct_0(X2))|~l1_struct_0(X2)), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_87, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_66]), c_0_67]), c_0_78]), c_0_69]), c_0_70]), c_0_71]), c_0_82])).
cnf(c_0_88, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_89, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_90, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_91, negated_conjecture, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_56])])).
cnf(c_0_92, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_93, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_56]), c_0_53])).
cnf(c_0_94, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_93]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 95
# Proof object clause steps            : 62
# Proof object formula steps           : 33
# Proof object conjectures             : 36
# Proof object clause conjectures      : 33
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 29
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 23
# Proof object simplifying inferences  : 82
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 38
# Processed clauses                    : 1840
# ...of these trivial                  : 4
# ...subsumed                          : 1456
# ...remaining for further processing  : 380
# Other redundant clauses eliminated   : 91
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 12
# Backward-rewritten                   : 98
# Generated clauses                    : 4526
# ...of the previous two non-trivial   : 4059
# Contextual simplify-reflections      : 14
# Paramodulations                      : 4410
# Factorizations                       : 10
# Equation resolutions                 : 106
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 236
#    Positive orientable unit clauses  : 15
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 212
# Current number of unprocessed clauses: 2108
# ...number of literals in the above   : 11289
# Current number of archived formulas  : 0
# Current number of archived clauses   : 146
# Clause-clause subsumption calls (NU) : 72028
# Rec. Clause-clause subsumption calls : 10654
# Non-unit clause-clause subsumptions  : 1291
# Unit Clause-clause subsumption calls : 643
# Rewrite failures with RHS unbound    : 69
# BW rewrite match attempts            : 61
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 67517

# -------------------------------------------------
# User time                : 0.074 s
# System time              : 0.003 s
# Total time               : 0.078 s
# Maximum resident set size: 28728 pages
