# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t28_connsp_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>~((![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X4,X3)<=>(v3_pre_topc(X4,X1)&(v4_pre_topc(X4,X1)&r2_hidden(X2,X4)))))&X3=k1_xboole_0))))), file('data/MPTP/proofs/342bbe10eb6342d1a968ba7785ff51cc.E_input', t28_connsp_2)).
fof(rc3_tops_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>?[X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))&(~(v1_xboole_0(X2))&(v3_pre_topc(X2,X1)&v4_pre_topc(X2,X1))))), file('data/MPTP/proofs/342bbe10eb6342d1a968ba7785ff51cc.E_input', rc3_tops_1)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/342bbe10eb6342d1a968ba7785ff51cc.E_input', dt_l1_pre_topc)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/342bbe10eb6342d1a968ba7785ff51cc.E_input', t5_subset)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/342bbe10eb6342d1a968ba7785ff51cc.E_input', t2_subset)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/342bbe10eb6342d1a968ba7785ff51cc.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/342bbe10eb6342d1a968ba7785ff51cc.E_input', d4_subset_1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/342bbe10eb6342d1a968ba7785ff51cc.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/342bbe10eb6342d1a968ba7785ff51cc.E_input', t2_xboole_1)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/342bbe10eb6342d1a968ba7785ff51cc.E_input', d3_struct_0)).
fof(fc4_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>v4_pre_topc(k2_struct_0(X1),X1)), file('data/MPTP/proofs/342bbe10eb6342d1a968ba7785ff51cc.E_input', fc4_pre_topc)).
fof(t9_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k4_tarski(X3,X4))))))), file('data/MPTP/proofs/342bbe10eb6342d1a968ba7785ff51cc.E_input', t9_mcart_1)).
fof(d5_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>r2_hidden(X2,u1_pre_topc(X1))))), file('data/MPTP/proofs/342bbe10eb6342d1a968ba7785ff51cc.E_input', d5_pre_topc)).
fof(d1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1))))))))), file('data/MPTP/proofs/342bbe10eb6342d1a968ba7785ff51cc.E_input', d1_pre_topc)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/342bbe10eb6342d1a968ba7785ff51cc.E_input', fc1_xboole_0)).
fof(c_0_15, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>~((![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X4,X3)<=>(v3_pre_topc(X4,X1)&(v4_pre_topc(X4,X1)&r2_hidden(X2,X4)))))&X3=k1_xboole_0)))))), inference(assume_negation,[status(cth)],[t28_connsp_2])).
fof(c_0_16, plain, ![X10]:((m1_subset_1(esk4_1(X10),k1_zfmisc_1(u1_struct_0(X10)))|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))&((~v1_xboole_0(esk4_1(X10))|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))&((v3_pre_topc(esk4_1(X10),X10)|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))&(v4_pre_topc(esk4_1(X10),X10)|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[rc3_tops_1])])])])])).
fof(c_0_17, negated_conjecture, ![X8]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))&((((v3_pre_topc(X8,esk1_0)|~r2_hidden(X8,esk3_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(esk1_0))))&((v4_pre_topc(X8,esk1_0)|~r2_hidden(X8,esk3_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(esk1_0))))&(r2_hidden(esk2_0,X8)|~r2_hidden(X8,esk3_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(esk1_0))))))&(~v3_pre_topc(X8,esk1_0)|(~v4_pre_topc(X8,esk1_0)|~r2_hidden(esk2_0,X8))|r2_hidden(X8,esk3_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(esk1_0)))))&esk3_0=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_15])])])])])])).
fof(c_0_18, plain, ![X24]:(~l1_pre_topc(X24)|l1_struct_0(X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_19, plain, ![X27, X28, X29]:(~r2_hidden(X27,X28)|(~m1_subset_1(X28,k1_zfmisc_1(X29))|~v1_xboole_0(X29))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
cnf(c_0_20, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_24, plain, ![X25, X26]:(~m1_subset_1(X25,X26)|(v1_xboole_0(X26)|r2_hidden(X25,X26))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_25, plain, ![X9]:m1_subset_1(k2_subset_1(X9),k1_zfmisc_1(X9)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_26, plain, ![X17]:k2_subset_1(X17)=X17, inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_27, plain, ![X22, X23]:(~r2_hidden(X22,X23)|~r1_tarski(X23,X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_28, plain, ![X12]:r1_tarski(k1_xboole_0,X12), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_29, plain, ![X13]:(~l1_struct_0(X13)|k2_struct_0(X13)=u1_struct_0(X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_30, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_31, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_32, negated_conjecture, (m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])]), c_0_23])).
cnf(c_0_33, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_35, negated_conjecture, (r2_hidden(X1,esk3_0)|~v3_pre_topc(X1,esk1_0)|~v4_pre_topc(X1,esk1_0)|~r2_hidden(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_36, negated_conjecture, (esk3_0=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_37, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_38, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_39, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_40, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_41, plain, ![X16]:(~v2_pre_topc(X16)|~l1_pre_topc(X16)|v4_pre_topc(k2_struct_0(X16),X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc4_pre_topc])])).
cnf(c_0_42, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_43, negated_conjecture, (l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_30, c_0_21])).
cnf(c_0_44, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))|~r2_hidden(X1,esk4_1(esk1_0))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_45, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
fof(c_0_46, plain, ![X18, X20, X21]:((r2_hidden(esk5_1(X18),X18)|X18=k1_xboole_0)&((~r2_hidden(X20,X18)|esk5_1(X18)!=k4_tarski(X20,X21)|X18=k1_xboole_0)&(~r2_hidden(X21,X18)|esk5_1(X18)!=k4_tarski(X20,X21)|X18=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_mcart_1])])])])])])).
cnf(c_0_47, negated_conjecture, (r2_hidden(X1,k1_xboole_0)|~v4_pre_topc(X1,esk1_0)|~v3_pre_topc(X1,esk1_0)|~r2_hidden(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(rw,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_48, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_49, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_50, plain, (v4_pre_topc(k2_struct_0(X1),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_51, negated_conjecture, (k2_struct_0(esk1_0)=u1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(X1,esk4_1(esk1_0))), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_53, plain, (r2_hidden(esk5_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_54, plain, ![X14, X15]:((~v3_pre_topc(X15,X14)|r2_hidden(X15,u1_pre_topc(X14))|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X14)))|~l1_pre_topc(X14))&(~r2_hidden(X15,u1_pre_topc(X14))|v3_pre_topc(X15,X14)|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X14)))|~l1_pre_topc(X14))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_pre_topc])])])])).
fof(c_0_55, plain, ![X30, X31, X32, X33]:(((r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|~v2_pre_topc(X30)|~l1_pre_topc(X30))&((~m1_subset_1(X31,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X30))))|(~r1_tarski(X31,u1_pre_topc(X30))|r2_hidden(k5_setfam_1(u1_struct_0(X30),X31),u1_pre_topc(X30)))|~v2_pre_topc(X30)|~l1_pre_topc(X30))&(~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X30)))|(~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X30)))|(~r2_hidden(X32,u1_pre_topc(X30))|~r2_hidden(X33,u1_pre_topc(X30))|r2_hidden(k9_subset_1(u1_struct_0(X30),X32,X33),u1_pre_topc(X30))))|~v2_pre_topc(X30)|~l1_pre_topc(X30))))&(((m1_subset_1(esk7_1(X30),k1_zfmisc_1(u1_struct_0(X30)))|m1_subset_1(esk6_1(X30),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X30))))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30))&((m1_subset_1(esk8_1(X30),k1_zfmisc_1(u1_struct_0(X30)))|m1_subset_1(esk6_1(X30),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X30))))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30))&(((r2_hidden(esk7_1(X30),u1_pre_topc(X30))|m1_subset_1(esk6_1(X30),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X30))))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30))&(r2_hidden(esk8_1(X30),u1_pre_topc(X30))|m1_subset_1(esk6_1(X30),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X30))))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30)))&(~r2_hidden(k9_subset_1(u1_struct_0(X30),esk7_1(X30),esk8_1(X30)),u1_pre_topc(X30))|m1_subset_1(esk6_1(X30),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X30))))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30)))))&(((m1_subset_1(esk7_1(X30),k1_zfmisc_1(u1_struct_0(X30)))|r1_tarski(esk6_1(X30),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30))&((m1_subset_1(esk8_1(X30),k1_zfmisc_1(u1_struct_0(X30)))|r1_tarski(esk6_1(X30),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30))&(((r2_hidden(esk7_1(X30),u1_pre_topc(X30))|r1_tarski(esk6_1(X30),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30))&(r2_hidden(esk8_1(X30),u1_pre_topc(X30))|r1_tarski(esk6_1(X30),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30)))&(~r2_hidden(k9_subset_1(u1_struct_0(X30),esk7_1(X30),esk8_1(X30)),u1_pre_topc(X30))|r1_tarski(esk6_1(X30),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30)))))&((m1_subset_1(esk7_1(X30),k1_zfmisc_1(u1_struct_0(X30)))|~r2_hidden(k5_setfam_1(u1_struct_0(X30),esk6_1(X30)),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30))&((m1_subset_1(esk8_1(X30),k1_zfmisc_1(u1_struct_0(X30)))|~r2_hidden(k5_setfam_1(u1_struct_0(X30),esk6_1(X30)),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30))&(((r2_hidden(esk7_1(X30),u1_pre_topc(X30))|~r2_hidden(k5_setfam_1(u1_struct_0(X30),esk6_1(X30)),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30))&(r2_hidden(esk8_1(X30),u1_pre_topc(X30))|~r2_hidden(k5_setfam_1(u1_struct_0(X30),esk6_1(X30)),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30)))&(~r2_hidden(k9_subset_1(u1_struct_0(X30),esk7_1(X30),esk8_1(X30)),u1_pre_topc(X30))|~r2_hidden(k5_setfam_1(u1_struct_0(X30),esk6_1(X30)),u1_pre_topc(X30))|~r2_hidden(u1_struct_0(X30),u1_pre_topc(X30))|v2_pre_topc(X30)|~l1_pre_topc(X30)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])).
cnf(c_0_56, negated_conjecture, (~v4_pre_topc(u1_struct_0(esk1_0),esk1_0)|~v3_pre_topc(u1_struct_0(esk1_0),esk1_0)|~r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])).
cnf(c_0_57, negated_conjecture, (v4_pre_topc(u1_struct_0(esk1_0),esk1_0)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_21]), c_0_22])]), c_0_51])).
cnf(c_0_58, plain, (v2_struct_0(X1)|~v1_xboole_0(esk4_1(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_59, negated_conjecture, (esk4_1(esk1_0)=k1_xboole_0|r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_60, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_61, plain, (v3_pre_topc(X1,X2)|~r2_hidden(X1,u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_62, plain, (r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_63, negated_conjecture, (~v3_pre_topc(u1_struct_0(esk1_0),esk1_0)|~r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_57])])).
cnf(c_0_64, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_60]), c_0_21]), c_0_22])]), c_0_23])).
cnf(c_0_65, plain, (v3_pre_topc(u1_struct_0(X1),X1)|~r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_61, c_0_48])).
cnf(c_0_66, negated_conjecture, (r2_hidden(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_21]), c_0_22])])).
cnf(c_0_67, negated_conjecture, (~v3_pre_topc(u1_struct_0(esk1_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_64])])).
cnf(c_0_68, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_21])]), c_0_67]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 69
# Proof object clause steps            : 39
# Proof object formula steps           : 30
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 23
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 47
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 46
# Processed clauses                    : 157
# ...of these trivial                  : 0
# ...subsumed                          : 13
# ...remaining for further processing  : 144
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 6
# Generated clauses                    : 90
# ...of the previous two non-trivial   : 78
# Contextual simplify-reflections      : 2
# Paramodulations                      : 90
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 91
#    Positive orientable unit clauses  : 20
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 65
# Current number of unprocessed clauses: 7
# ...number of literals in the above   : 22
# Current number of archived formulas  : 0
# Current number of archived clauses   : 54
# Clause-clause subsumption calls (NU) : 1596
# Rec. Clause-clause subsumption calls : 299
# Non-unit clause-clause subsumptions  : 7
# Unit Clause-clause subsumption calls : 84
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5789

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.003 s
# Total time               : 0.013 s
# Maximum resident set size: 28620 pages
