# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.011 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/3459d4b40eb4450589d1d79b3a4b3142.E_input', d1_tarski)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/3459d4b40eb4450589d1d79b3a4b3142.E_input', t69_enumset1)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/3459d4b40eb4450589d1d79b3a4b3142.E_input', redefinition_k6_domain_1)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/3459d4b40eb4450589d1d79b3a4b3142.E_input', d9_lattice3)).
fof(t39_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2))), file('data/MPTP/proofs/3459d4b40eb4450589d1d79b3a4b3142.E_input', t39_yellow_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/3459d4b40eb4450589d1d79b3a4b3142.E_input', t24_orders_2)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/3459d4b40eb4450589d1d79b3a4b3142.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/3459d4b40eb4450589d1d79b3a4b3142.E_input', dt_l1_orders_2)).
fof(t18_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))&((r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))=>(X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3))))))))), file('data/MPTP/proofs/3459d4b40eb4450589d1d79b3a4b3142.E_input', t18_yellow_0)).
fof(d7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X5)=>r1_orders_2(X1,X4,X5))))=>X4=X3))))))), file('data/MPTP/proofs/3459d4b40eb4450589d1d79b3a4b3142.E_input', d7_yellow_0)).
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/3459d4b40eb4450589d1d79b3a4b3142.E_input', d9_yellow_0)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/3459d4b40eb4450589d1d79b3a4b3142.E_input', dt_k1_yellow_0)).
fof(t7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2)))))))), file('data/MPTP/proofs/3459d4b40eb4450589d1d79b3a4b3142.E_input', t7_yellow_0)).
fof(d10_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_yellow_0(X1,X2)=>(X3=k2_yellow_0(X1,X2)<=>(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))))))), file('data/MPTP/proofs/3459d4b40eb4450589d1d79b3a4b3142.E_input', d10_yellow_0)).
fof(dt_k2_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/3459d4b40eb4450589d1d79b3a4b3142.E_input', dt_k2_yellow_0)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/3459d4b40eb4450589d1d79b3a4b3142.E_input', d2_tarski)).
fof(t30_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4)))))&((r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))=>(X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3)))))), file('data/MPTP/proofs/3459d4b40eb4450589d1d79b3a4b3142.E_input', t30_yellow_0)).
fof(t19_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))))), file('data/MPTP/proofs/3459d4b40eb4450589d1d79b3a4b3142.E_input', t19_yellow_0)).
fof(c_0_18, plain, ![X45, X46, X47, X48, X49, X50]:(((~r2_hidden(X47,X46)|X47=X45|X46!=k1_tarski(X45))&(X48!=X45|r2_hidden(X48,X46)|X46!=k1_tarski(X45)))&((~r2_hidden(esk7_2(X49,X50),X50)|esk7_2(X49,X50)!=X49|X50=k1_tarski(X49))&(r2_hidden(esk7_2(X49,X50),X50)|esk7_2(X49,X50)=X49|X50=k1_tarski(X49)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_19, plain, ![X75]:k2_tarski(X75,X75)=k1_tarski(X75), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_20, plain, ![X53, X54]:(v1_xboole_0(X53)|~m1_subset_1(X54,X53)|k6_domain_1(X53,X54)=k1_tarski(X54)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_21, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_24, plain, ![X67, X68, X69, X70]:((~r2_lattice3(X67,X68,X69)|(~m1_subset_1(X70,u1_struct_0(X67))|(~r2_hidden(X70,X68)|r1_orders_2(X67,X70,X69)))|~m1_subset_1(X69,u1_struct_0(X67))|~l1_orders_2(X67))&((m1_subset_1(esk12_3(X67,X68,X69),u1_struct_0(X67))|r2_lattice3(X67,X68,X69)|~m1_subset_1(X69,u1_struct_0(X67))|~l1_orders_2(X67))&((r2_hidden(esk12_3(X67,X68,X69),X68)|r2_lattice3(X67,X68,X69)|~m1_subset_1(X69,u1_struct_0(X67))|~l1_orders_2(X67))&(~r1_orders_2(X67,esk12_3(X67,X68,X69),X69)|r2_lattice3(X67,X68,X69)|~m1_subset_1(X69,u1_struct_0(X67))|~l1_orders_2(X67))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
cnf(c_0_25, plain, (X1=X3|X2!=k2_tarski(X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_21, c_0_22])).
fof(c_0_26, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))&((r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))=>(X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3))))))), introduced(definition)).
cnf(c_0_27, plain, (k2_tarski(X2,X2)=k6_domain_1(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_23, c_0_22])).
cnf(c_0_28, plain, (m1_subset_1(esk12_3(X1,X2,X3),u1_struct_0(X1))|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29, plain, (X1=X2|~r2_hidden(X1,k2_tarski(X2,X2))), inference(er,[status(thm)],[c_0_25])).
cnf(c_0_30, plain, (r2_hidden(esk12_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_31, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2)))), inference(assume_negation,[status(cth)],[t39_yellow_0])).
fof(c_0_32, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))&((r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))=>(X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3))))))), inference(split_equiv,[status(thm)],[c_0_26])).
cnf(c_0_33, plain, (k2_tarski(esk12_3(X1,X2,X3),esk12_3(X1,X2,X3))=k6_domain_1(u1_struct_0(X1),esk12_3(X1,X2,X3))|v1_xboole_0(u1_struct_0(X1))|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34, plain, (esk12_3(X1,k2_tarski(X2,X2),X3)=X2|r2_lattice3(X1,k2_tarski(X2,X2),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
fof(c_0_35, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_31])])])])).
fof(c_0_36, plain, ![X76, X77, X78, X79, X80]:(((r1_orders_2(X78,X77,X79)|(X79!=k10_lattice3(X78,X77,X76)|~r1_yellow_0(X78,k2_tarski(X77,X76)))|~m1_subset_1(X79,u1_struct_0(X78))|~epred1_3(X78,X77,X76))&((r1_orders_2(X78,X76,X79)|(X79!=k10_lattice3(X78,X77,X76)|~r1_yellow_0(X78,k2_tarski(X77,X76)))|~m1_subset_1(X79,u1_struct_0(X78))|~epred1_3(X78,X77,X76))&(~m1_subset_1(X80,u1_struct_0(X78))|(~r1_orders_2(X78,X77,X80)|~r1_orders_2(X78,X76,X80)|r1_orders_2(X78,X79,X80))|(X79!=k10_lattice3(X78,X77,X76)|~r1_yellow_0(X78,k2_tarski(X77,X76)))|~m1_subset_1(X79,u1_struct_0(X78))|~epred1_3(X78,X77,X76))))&(((X79=k10_lattice3(X78,X77,X76)|(m1_subset_1(esk13_4(X76,X77,X78,X79),u1_struct_0(X78))|~r1_orders_2(X78,X76,X79)|~r1_orders_2(X78,X77,X79))|~m1_subset_1(X79,u1_struct_0(X78))|~epred1_3(X78,X77,X76))&(r1_yellow_0(X78,k2_tarski(X77,X76))|(m1_subset_1(esk13_4(X76,X77,X78,X79),u1_struct_0(X78))|~r1_orders_2(X78,X76,X79)|~r1_orders_2(X78,X77,X79))|~m1_subset_1(X79,u1_struct_0(X78))|~epred1_3(X78,X77,X76)))&((((X79=k10_lattice3(X78,X77,X76)|(r1_orders_2(X78,X77,esk13_4(X76,X77,X78,X79))|~r1_orders_2(X78,X76,X79)|~r1_orders_2(X78,X77,X79))|~m1_subset_1(X79,u1_struct_0(X78))|~epred1_3(X78,X77,X76))&(r1_yellow_0(X78,k2_tarski(X77,X76))|(r1_orders_2(X78,X77,esk13_4(X76,X77,X78,X79))|~r1_orders_2(X78,X76,X79)|~r1_orders_2(X78,X77,X79))|~m1_subset_1(X79,u1_struct_0(X78))|~epred1_3(X78,X77,X76)))&((X79=k10_lattice3(X78,X77,X76)|(r1_orders_2(X78,X76,esk13_4(X76,X77,X78,X79))|~r1_orders_2(X78,X76,X79)|~r1_orders_2(X78,X77,X79))|~m1_subset_1(X79,u1_struct_0(X78))|~epred1_3(X78,X77,X76))&(r1_yellow_0(X78,k2_tarski(X77,X76))|(r1_orders_2(X78,X76,esk13_4(X76,X77,X78,X79))|~r1_orders_2(X78,X76,X79)|~r1_orders_2(X78,X77,X79))|~m1_subset_1(X79,u1_struct_0(X78))|~epred1_3(X78,X77,X76))))&((X79=k10_lattice3(X78,X77,X76)|(~r1_orders_2(X78,X79,esk13_4(X76,X77,X78,X79))|~r1_orders_2(X78,X76,X79)|~r1_orders_2(X78,X77,X79))|~m1_subset_1(X79,u1_struct_0(X78))|~epred1_3(X78,X77,X76))&(r1_yellow_0(X78,k2_tarski(X77,X76))|(~r1_orders_2(X78,X79,esk13_4(X76,X77,X78,X79))|~r1_orders_2(X78,X76,X79)|~r1_orders_2(X78,X77,X79))|~m1_subset_1(X79,u1_struct_0(X78))|~epred1_3(X78,X77,X76)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_32])])])])])).
fof(c_0_37, plain, ![X19, X20]:(v2_struct_0(X19)|(~v3_orders_2(X19)|~l1_orders_2(X19))|(~m1_subset_1(X20,u1_struct_0(X19))|r1_orders_2(X19,X20,X20))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_38, plain, ![X55]:(v2_struct_0(X55)|~l1_struct_0(X55)|~v1_xboole_0(u1_struct_0(X55))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_39, plain, (k2_tarski(X1,X1)=k6_domain_1(u1_struct_0(X2),X1)|v1_xboole_0(u1_struct_0(X2))|r2_lattice3(X2,k2_tarski(X1,X1),X3)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_40, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_42, plain, (X1=k10_lattice3(X2,X3,X4)|~r1_orders_2(X2,X1,esk13_4(X4,X3,X2,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_43, plain, (X1=k10_lattice3(X2,X3,X4)|r1_orders_2(X2,X4,esk13_4(X4,X3,X2,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_45, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_46, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_47, plain, (X1=k10_lattice3(X2,X3,X4)|r1_orders_2(X2,X3,esk13_4(X4,X3,X2,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_48, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_49, negated_conjecture, (k2_tarski(X1,X1)=k6_domain_1(u1_struct_0(esk1_0),X1)|v1_xboole_0(u1_struct_0(esk1_0))|r2_lattice3(esk1_0,k2_tarski(X1,X1),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])])).
fof(c_0_50, plain, ![X52]:(~l1_orders_2(X52)|l1_struct_0(X52)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_51, plain, (k10_lattice3(X1,X2,X3)=X3|~epred1_3(X1,X2,X3)|~r1_orders_2(X1,X3,X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_52, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_41])]), c_0_46])).
cnf(c_0_53, plain, (k10_lattice3(X1,X2,X3)=X2|~epred1_3(X1,X2,X3)|~r1_orders_2(X1,X3,X2)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_42, c_0_47])).
fof(c_0_54, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>epred1_3(X1,X2,X3)))), inference(apply_def,[status(thm)],[t18_yellow_0, c_0_26])).
cnf(c_0_55, plain, (r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk12_3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_56, negated_conjecture, (k2_tarski(X1,X1)=k6_domain_1(u1_struct_0(esk1_0),X1)|r2_lattice3(esk1_0,k2_tarski(X1,X1),esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_46])).
cnf(c_0_57, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_58, negated_conjecture, (k10_lattice3(esk1_0,X1,X2)=X2|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_59, negated_conjecture, (k10_lattice3(esk1_0,X1,X2)=X1|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_53, c_0_52])).
fof(c_0_60, plain, ![X21, X22, X23]:(~v5_orders_2(X21)|~l1_orders_2(X21)|(~m1_subset_1(X22,u1_struct_0(X21))|(~m1_subset_1(X23,u1_struct_0(X21))|epred1_3(X21,X22,X23)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_54])])])).
cnf(c_0_61, plain, (r2_lattice3(X1,k2_tarski(X2,X2),X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_55, c_0_34])).
cnf(c_0_62, negated_conjecture, (k2_tarski(X1,X1)=k6_domain_1(u1_struct_0(esk1_0),X1)|r2_lattice3(esk1_0,k2_tarski(X1,X1),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_41])])).
cnf(c_0_63, negated_conjecture, (X1=X2|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_64, plain, (epred1_3(X1,X2,X3)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_65, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_66, plain, ![X56, X57, X59, X60, X62, X63, X66]:(((m1_subset_1(esk8_2(X56,X57),u1_struct_0(X56))|~r1_yellow_0(X56,X57)|~l1_orders_2(X56))&((r2_lattice3(X56,X57,esk8_2(X56,X57))|~r1_yellow_0(X56,X57)|~l1_orders_2(X56))&((~m1_subset_1(X59,u1_struct_0(X56))|(~r2_lattice3(X56,X57,X59)|r1_orders_2(X56,esk8_2(X56,X57),X59))|~r1_yellow_0(X56,X57)|~l1_orders_2(X56))&((m1_subset_1(esk9_3(X56,X57,X60),u1_struct_0(X56))|~r2_lattice3(X56,X57,X60)|X60=esk8_2(X56,X57)|~m1_subset_1(X60,u1_struct_0(X56))|~r1_yellow_0(X56,X57)|~l1_orders_2(X56))&((r2_lattice3(X56,X57,esk9_3(X56,X57,X60))|~r2_lattice3(X56,X57,X60)|X60=esk8_2(X56,X57)|~m1_subset_1(X60,u1_struct_0(X56))|~r1_yellow_0(X56,X57)|~l1_orders_2(X56))&(~r1_orders_2(X56,X60,esk9_3(X56,X57,X60))|~r2_lattice3(X56,X57,X60)|X60=esk8_2(X56,X57)|~m1_subset_1(X60,u1_struct_0(X56))|~r1_yellow_0(X56,X57)|~l1_orders_2(X56)))))))&(((m1_subset_1(esk11_3(X56,X62,X63),u1_struct_0(X56))|m1_subset_1(esk10_3(X56,X62,X63),u1_struct_0(X56))|~r2_lattice3(X56,X62,X63)|~m1_subset_1(X63,u1_struct_0(X56))|r1_yellow_0(X56,X62)|~l1_orders_2(X56))&(((r2_lattice3(X56,X62,esk11_3(X56,X62,X63))|m1_subset_1(esk10_3(X56,X62,X63),u1_struct_0(X56))|~r2_lattice3(X56,X62,X63)|~m1_subset_1(X63,u1_struct_0(X56))|r1_yellow_0(X56,X62)|~l1_orders_2(X56))&(~m1_subset_1(X66,u1_struct_0(X56))|(~r2_lattice3(X56,X62,X66)|r1_orders_2(X56,esk11_3(X56,X62,X63),X66))|m1_subset_1(esk10_3(X56,X62,X63),u1_struct_0(X56))|~r2_lattice3(X56,X62,X63)|~m1_subset_1(X63,u1_struct_0(X56))|r1_yellow_0(X56,X62)|~l1_orders_2(X56)))&(esk11_3(X56,X62,X63)!=X63|m1_subset_1(esk10_3(X56,X62,X63),u1_struct_0(X56))|~r2_lattice3(X56,X62,X63)|~m1_subset_1(X63,u1_struct_0(X56))|r1_yellow_0(X56,X62)|~l1_orders_2(X56))))&(((m1_subset_1(esk11_3(X56,X62,X63),u1_struct_0(X56))|r2_lattice3(X56,X62,esk10_3(X56,X62,X63))|~r2_lattice3(X56,X62,X63)|~m1_subset_1(X63,u1_struct_0(X56))|r1_yellow_0(X56,X62)|~l1_orders_2(X56))&(((r2_lattice3(X56,X62,esk11_3(X56,X62,X63))|r2_lattice3(X56,X62,esk10_3(X56,X62,X63))|~r2_lattice3(X56,X62,X63)|~m1_subset_1(X63,u1_struct_0(X56))|r1_yellow_0(X56,X62)|~l1_orders_2(X56))&(~m1_subset_1(X66,u1_struct_0(X56))|(~r2_lattice3(X56,X62,X66)|r1_orders_2(X56,esk11_3(X56,X62,X63),X66))|r2_lattice3(X56,X62,esk10_3(X56,X62,X63))|~r2_lattice3(X56,X62,X63)|~m1_subset_1(X63,u1_struct_0(X56))|r1_yellow_0(X56,X62)|~l1_orders_2(X56)))&(esk11_3(X56,X62,X63)!=X63|r2_lattice3(X56,X62,esk10_3(X56,X62,X63))|~r2_lattice3(X56,X62,X63)|~m1_subset_1(X63,u1_struct_0(X56))|r1_yellow_0(X56,X62)|~l1_orders_2(X56))))&((m1_subset_1(esk11_3(X56,X62,X63),u1_struct_0(X56))|~r1_orders_2(X56,X63,esk10_3(X56,X62,X63))|~r2_lattice3(X56,X62,X63)|~m1_subset_1(X63,u1_struct_0(X56))|r1_yellow_0(X56,X62)|~l1_orders_2(X56))&(((r2_lattice3(X56,X62,esk11_3(X56,X62,X63))|~r1_orders_2(X56,X63,esk10_3(X56,X62,X63))|~r2_lattice3(X56,X62,X63)|~m1_subset_1(X63,u1_struct_0(X56))|r1_yellow_0(X56,X62)|~l1_orders_2(X56))&(~m1_subset_1(X66,u1_struct_0(X56))|(~r2_lattice3(X56,X62,X66)|r1_orders_2(X56,esk11_3(X56,X62,X63),X66))|~r1_orders_2(X56,X63,esk10_3(X56,X62,X63))|~r2_lattice3(X56,X62,X63)|~m1_subset_1(X63,u1_struct_0(X56))|r1_yellow_0(X56,X62)|~l1_orders_2(X56)))&(esk11_3(X56,X62,X63)!=X63|~r1_orders_2(X56,X63,esk10_3(X56,X62,X63))|~r2_lattice3(X56,X62,X63)|~m1_subset_1(X63,u1_struct_0(X56))|r1_yellow_0(X56,X62)|~l1_orders_2(X56))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_yellow_0])])])])])])).
fof(c_0_67, plain, ![X14, X15, X16, X17]:(((r2_lattice3(X14,X15,X16)|X16!=k1_yellow_0(X14,X15)|~r1_yellow_0(X14,X15)|~m1_subset_1(X16,u1_struct_0(X14))|~l1_orders_2(X14))&(~m1_subset_1(X17,u1_struct_0(X14))|(~r2_lattice3(X14,X15,X17)|r1_orders_2(X14,X16,X17))|X16!=k1_yellow_0(X14,X15)|~r1_yellow_0(X14,X15)|~m1_subset_1(X16,u1_struct_0(X14))|~l1_orders_2(X14)))&((m1_subset_1(esk4_3(X14,X15,X16),u1_struct_0(X14))|~r2_lattice3(X14,X15,X16)|X16=k1_yellow_0(X14,X15)|~r1_yellow_0(X14,X15)|~m1_subset_1(X16,u1_struct_0(X14))|~l1_orders_2(X14))&((r2_lattice3(X14,X15,esk4_3(X14,X15,X16))|~r2_lattice3(X14,X15,X16)|X16=k1_yellow_0(X14,X15)|~r1_yellow_0(X14,X15)|~m1_subset_1(X16,u1_struct_0(X14))|~l1_orders_2(X14))&(~r1_orders_2(X14,X16,esk4_3(X14,X15,X16))|~r2_lattice3(X14,X15,X16)|X16=k1_yellow_0(X14,X15)|~r1_yellow_0(X14,X15)|~m1_subset_1(X16,u1_struct_0(X14))|~l1_orders_2(X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
fof(c_0_68, plain, ![X40, X41]:(~l1_orders_2(X40)|m1_subset_1(k1_yellow_0(X40,X41),u1_struct_0(X40))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
cnf(c_0_69, negated_conjecture, (r2_lattice3(X1,k6_domain_1(u1_struct_0(esk1_0),X2),X3)|r2_lattice3(esk1_0,k2_tarski(X2,X2),esk2_0)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_70, negated_conjecture, (X1=X2|~r1_orders_2(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_41]), c_0_65])])).
cnf(c_0_71, plain, (r1_orders_2(X2,esk8_2(X2,X3),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r1_yellow_0(X2,X3)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_72, plain, (r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~r1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_73, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
fof(c_0_74, plain, ![X42, X43, X44]:((~r1_lattice3(X42,k1_tarski(X44),X43)|r1_orders_2(X42,X43,X44)|~m1_subset_1(X44,u1_struct_0(X42))|~m1_subset_1(X43,u1_struct_0(X42))|~l1_orders_2(X42))&((~r1_orders_2(X42,X43,X44)|r1_lattice3(X42,k1_tarski(X44),X43)|~m1_subset_1(X44,u1_struct_0(X42))|~m1_subset_1(X43,u1_struct_0(X42))|~l1_orders_2(X42))&((~r2_lattice3(X42,k1_tarski(X44),X43)|r1_orders_2(X42,X44,X43)|~m1_subset_1(X44,u1_struct_0(X42))|~m1_subset_1(X43,u1_struct_0(X42))|~l1_orders_2(X42))&(~r1_orders_2(X42,X44,X43)|r2_lattice3(X42,k1_tarski(X44),X43)|~m1_subset_1(X44,u1_struct_0(X42))|~m1_subset_1(X43,u1_struct_0(X42))|~l1_orders_2(X42))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])).
cnf(c_0_75, negated_conjecture, (r2_lattice3(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1),esk2_0)|r2_lattice3(esk1_0,k2_tarski(X1,X1),esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_40]), c_0_41])])).
cnf(c_0_76, negated_conjecture, (esk8_2(esk1_0,X1)=X2|~r1_orders_2(esk1_0,X2,esk8_2(esk1_0,X1))|~r2_lattice3(esk1_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk8_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_41])])).
cnf(c_0_77, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),X3)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_72]), c_0_73])).
cnf(c_0_78, plain, (r1_orders_2(X1,X2,X3)|~r2_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_79, negated_conjecture, (r2_lattice3(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk2_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_52]), c_0_40])])).
cnf(c_0_80, negated_conjecture, (esk8_2(esk1_0,X1)=k1_yellow_0(esk1_0,X2)|~r2_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,X2))|~r2_lattice3(esk1_0,X2,esk8_2(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(esk8_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_41])])).
cnf(c_0_81, plain, (r2_lattice3(X1,X2,esk8_2(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_82, plain, (r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r2_lattice3(X1,k2_tarski(X2,X2),X3)), inference(rw,[status(thm)],[c_0_78, c_0_22])).
cnf(c_0_83, negated_conjecture, (r2_lattice3(esk1_0,k2_tarski(esk2_0,esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_79, c_0_62])).
cnf(c_0_84, negated_conjecture, (esk8_2(esk1_0,X1)=k1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk8_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_41])])).
cnf(c_0_85, plain, (m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_86, plain, (r2_lattice3(X1,X2,X3)|X3!=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_87, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_40]), c_0_41])])).
cnf(c_0_88, negated_conjecture, (esk8_2(esk1_0,X1)=k1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_41])])).
cnf(c_0_89, plain, (r2_lattice3(X1,X2,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_86]), c_0_73])).
fof(c_0_90, plain, ![X24, X25, X26, X27]:(((r1_lattice3(X24,X25,X26)|X26!=k2_yellow_0(X24,X25)|~r2_yellow_0(X24,X25)|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))&(~m1_subset_1(X27,u1_struct_0(X24))|(~r1_lattice3(X24,X25,X27)|r1_orders_2(X24,X27,X26))|X26!=k2_yellow_0(X24,X25)|~r2_yellow_0(X24,X25)|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24)))&((m1_subset_1(esk5_3(X24,X25,X26),u1_struct_0(X24))|~r1_lattice3(X24,X25,X26)|X26=k2_yellow_0(X24,X25)|~r2_yellow_0(X24,X25)|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))&((r1_lattice3(X24,X25,esk5_3(X24,X25,X26))|~r1_lattice3(X24,X25,X26)|X26=k2_yellow_0(X24,X25)|~r2_yellow_0(X24,X25)|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))&(~r1_orders_2(X24,esk5_3(X24,X25,X26),X26)|~r1_lattice3(X24,X25,X26)|X26=k2_yellow_0(X24,X25)|~r2_yellow_0(X24,X25)|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_yellow_0])])])])])).
fof(c_0_91, plain, ![X38, X39]:(~l1_orders_2(X38)|m1_subset_1(k2_yellow_0(X38,X39),u1_struct_0(X38))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_yellow_0])])).
cnf(c_0_92, plain, (k10_lattice3(esk1_0,esk2_0,X1)=esk2_0|~epred1_3(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_87]), c_0_40])])).
cnf(c_0_93, negated_conjecture, (esk8_2(esk1_0,X1)=k1_yellow_0(esk1_0,X1)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_41])])).
cnf(c_0_94, plain, (r1_orders_2(X2,X1,X4)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|X4!=k2_yellow_0(X2,X3)|~r2_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_95, plain, (m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_96, plain, (r1_orders_2(X1,X3,X2)|~r1_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_97, plain, (r1_lattice3(X1,X2,X3)|X3!=k2_yellow_0(X1,X2)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_98, negated_conjecture, (esk2_0=X1|~epred1_3(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_58, c_0_92])).
cnf(c_0_99, plain, (r1_orders_2(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_100, negated_conjecture, (r2_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_93]), c_0_41])])).
cnf(c_0_101, plain, (r1_orders_2(X1,X2,k2_yellow_0(X1,X3))|~r1_lattice3(X1,X3,X2)|~r2_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_94]), c_0_95])).
cnf(c_0_102, plain, (r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,k2_tarski(X2,X2),X3)), inference(rw,[status(thm)],[c_0_96, c_0_22])).
cnf(c_0_103, plain, (r1_lattice3(X1,X2,k2_yellow_0(X1,X2))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_97]), c_0_95])).
fof(c_0_104, plain, ![X29, X30, X31, X32, X33, X34, X35, X36]:(((~r2_hidden(X32,X31)|(X32=X29|X32=X30)|X31!=k2_tarski(X29,X30))&((X33!=X29|r2_hidden(X33,X31)|X31!=k2_tarski(X29,X30))&(X33!=X30|r2_hidden(X33,X31)|X31!=k2_tarski(X29,X30))))&(((esk6_3(X34,X35,X36)!=X34|~r2_hidden(esk6_3(X34,X35,X36),X36)|X36=k2_tarski(X34,X35))&(esk6_3(X34,X35,X36)!=X35|~r2_hidden(esk6_3(X34,X35,X36),X36)|X36=k2_tarski(X34,X35)))&(r2_hidden(esk6_3(X34,X35,X36),X36)|(esk6_3(X34,X35,X36)=X34|esk6_3(X34,X35,X36)=X35)|X36=k2_tarski(X34,X35)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_105, negated_conjecture, (esk2_0=X1|~r1_orders_2(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_64]), c_0_40]), c_0_41]), c_0_65])])).
cnf(c_0_106, negated_conjecture, (r1_orders_2(esk1_0,X1,k1_yellow_0(esk1_0,X2))|~r2_hidden(X1,X2)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(k1_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_41])])).
cnf(c_0_107, negated_conjecture, (X1=k2_yellow_0(esk1_0,X2)|~r1_lattice3(esk1_0,X2,X1)|~r2_yellow_0(esk1_0,X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X2),X1)|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_101]), c_0_41])])).
cnf(c_0_108, plain, (r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(X2,X2)),X2)|~r2_yellow_0(X1,k2_tarski(X2,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_95])).
cnf(c_0_109, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_110, negated_conjecture, (k1_yellow_0(esk1_0,X1)=esk2_0|~r2_hidden(esk2_0,X1)|~r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),esk2_0)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_40])])).
cnf(c_0_111, negated_conjecture, (k2_yellow_0(esk1_0,k2_tarski(X1,X1))=X1|~r1_lattice3(esk1_0,k2_tarski(X1,X1),X1)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(X1,X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_41])])).
cnf(c_0_112, plain, (r1_lattice3(X1,k1_tarski(X3),X2)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_113, plain, (r2_hidden(X1,k2_tarski(X1,X2))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_109])])).
cnf(c_0_114, plain, (k2_tarski(k1_yellow_0(X1,X2),k1_yellow_0(X1,X2))=k6_domain_1(u1_struct_0(X1),k1_yellow_0(X1,X2))|v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_27, c_0_73])).
cnf(c_0_115, negated_conjecture, (k1_yellow_0(esk1_0,X1)=esk2_0|~r2_hidden(esk2_0,X1)|~r2_lattice3(esk1_0,X1,esk2_0)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_77]), c_0_40]), c_0_41])])).
cnf(c_0_116, plain, (k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X2))=k6_domain_1(u1_struct_0(X1),k2_yellow_0(X1,X2))|v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_27, c_0_95])).
cnf(c_0_117, negated_conjecture, (k2_yellow_0(esk1_0,k2_tarski(X1,X1))=X1|~r1_lattice3(esk1_0,k2_tarski(X1,X1),X1)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_95]), c_0_41])])).
cnf(c_0_118, plain, (r1_lattice3(X1,k2_tarski(X3,X3),X2)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[c_0_112, c_0_22])).
cnf(c_0_119, negated_conjecture, (k1_yellow_0(esk1_0,X1)=X2|~r1_orders_2(esk1_0,X2,k1_yellow_0(esk1_0,X1))|~r2_lattice3(esk1_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_77]), c_0_41])])).
cnf(c_0_120, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k1_yellow_0(X1,X2),k6_domain_1(u1_struct_0(X1),k1_yellow_0(X1,X2)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_113, c_0_114])).
cnf(c_0_121, negated_conjecture, (k1_yellow_0(esk1_0,X1)=esk2_0|~r2_hidden(esk2_0,X1)|~r2_lattice3(esk1_0,X1,esk2_0)|~r1_yellow_0(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_73]), c_0_41])])).
cnf(c_0_122, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k2_yellow_0(X1,X2),k6_domain_1(u1_struct_0(X1),k2_yellow_0(X1,X2)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_113, c_0_116])).
cnf(c_0_123, negated_conjecture, (k2_yellow_0(esk1_0,k2_tarski(X1,X1))=X1|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_41])]), c_0_52])).
cnf(c_0_124, negated_conjecture, (k1_yellow_0(esk1_0,X1)=X2|~r2_hidden(X2,X1)|~r2_lattice3(esk1_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_119, c_0_106])).
cnf(c_0_125, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))|~r2_hidden(esk2_0,X1)|~r2_lattice3(esk1_0,X1,esk2_0)|~r1_yellow_0(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121]), c_0_41])])).
cnf(c_0_126, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(X1,k6_domain_1(u1_struct_0(esk1_0),X1))|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_41])])).
cnf(c_0_127, negated_conjecture, (k1_yellow_0(esk1_0,X1)=X2|~r2_hidden(X2,X1)|~r2_lattice3(esk1_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_73]), c_0_41])])).
cnf(c_0_128, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~r2_lattice3(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk2_0)|~r1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_40])])).
cnf(c_0_129, negated_conjecture, (k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(X1),k1_yellow_0(X1,X2)))=k1_yellow_0(X1,X2)|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(X1),k1_yellow_0(X1,X2)))|~m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(esk1_0))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_123, c_0_114])).
cnf(c_0_130, negated_conjecture, (k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))=esk2_0|v1_xboole_0(u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~r2_lattice3(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk2_0)|~r1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128]), c_0_40])])).
cnf(c_0_131, negated_conjecture, (k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_132, negated_conjecture, (k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))=esk2_0|v1_xboole_0(u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~r2_lattice3(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk2_0)|~r1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130]), c_0_40]), c_0_41])])).
cnf(c_0_133, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_27, c_0_40])).
fof(c_0_134, plain, ![X8, X9, X10, X11, X12]:(((r2_lattice3(X8,X10,X9)|(X9!=k1_yellow_0(X8,X10)|~r1_yellow_0(X8,X10))|~m1_subset_1(X9,u1_struct_0(X8))|(~v5_orders_2(X8)|~l1_orders_2(X8)))&(~m1_subset_1(X11,u1_struct_0(X8))|(~r2_lattice3(X8,X10,X11)|r1_orders_2(X8,X9,X11))|(X9!=k1_yellow_0(X8,X10)|~r1_yellow_0(X8,X10))|~m1_subset_1(X9,u1_struct_0(X8))|(~v5_orders_2(X8)|~l1_orders_2(X8))))&(((X9=k1_yellow_0(X8,X12)|(m1_subset_1(esk3_3(X8,X9,X12),u1_struct_0(X8))|~r2_lattice3(X8,X12,X9))|~m1_subset_1(X9,u1_struct_0(X8))|(~v5_orders_2(X8)|~l1_orders_2(X8)))&(r1_yellow_0(X8,X12)|(m1_subset_1(esk3_3(X8,X9,X12),u1_struct_0(X8))|~r2_lattice3(X8,X12,X9))|~m1_subset_1(X9,u1_struct_0(X8))|(~v5_orders_2(X8)|~l1_orders_2(X8))))&(((X9=k1_yellow_0(X8,X12)|(r2_lattice3(X8,X12,esk3_3(X8,X9,X12))|~r2_lattice3(X8,X12,X9))|~m1_subset_1(X9,u1_struct_0(X8))|(~v5_orders_2(X8)|~l1_orders_2(X8)))&(r1_yellow_0(X8,X12)|(r2_lattice3(X8,X12,esk3_3(X8,X9,X12))|~r2_lattice3(X8,X12,X9))|~m1_subset_1(X9,u1_struct_0(X8))|(~v5_orders_2(X8)|~l1_orders_2(X8))))&((X9=k1_yellow_0(X8,X12)|(~r1_orders_2(X8,X9,esk3_3(X8,X9,X12))|~r2_lattice3(X8,X12,X9))|~m1_subset_1(X9,u1_struct_0(X8))|(~v5_orders_2(X8)|~l1_orders_2(X8)))&(r1_yellow_0(X8,X12)|(~r1_orders_2(X8,X9,esk3_3(X8,X9,X12))|~r2_lattice3(X8,X12,X9))|~m1_subset_1(X9,u1_struct_0(X8))|(~v5_orders_2(X8)|~l1_orders_2(X8))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_yellow_0])])])])])])).
fof(c_0_135, plain, ![X3, X2, X1]:(epred2_3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))), introduced(definition)).
cnf(c_0_136, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~r2_lattice3(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk2_0)|~r1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_132]), c_0_130])).
cnf(c_0_137, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_133]), c_0_46])).
cnf(c_0_138, plain, (r1_yellow_0(X1,X2)|r2_lattice3(X1,X2,esk3_3(X1,X3,X2))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_139, plain, (r1_yellow_0(X1,X2)|m1_subset_1(esk3_3(X1,X3,X2),u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_134])).
fof(c_0_140, plain, ![X3, X2, X1]:(epred2_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))), inference(split_equiv,[status(thm)],[c_0_135])).
cnf(c_0_141, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~l1_struct_0(esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~r1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_137]), c_0_83])])).
cnf(c_0_142, plain, (r1_yellow_0(X1,X2)|~r1_orders_2(X1,X3,esk3_3(X1,X3,X2))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_143, plain, (r1_orders_2(X1,X2,esk3_3(X1,X3,X4))|r1_yellow_0(X1,X4)|~r2_hidden(X2,X4)|~r2_lattice3(X1,X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_138]), c_0_139])).
fof(c_0_144, plain, ![X82, X83, X84, X85, X86]:(((r1_orders_2(X84,X85,X83)|(X85!=k11_lattice3(X84,X83,X82)|~r2_yellow_0(X84,k2_tarski(X83,X82)))|~m1_subset_1(X85,u1_struct_0(X84))|~epred2_3(X84,X83,X82))&((r1_orders_2(X84,X85,X82)|(X85!=k11_lattice3(X84,X83,X82)|~r2_yellow_0(X84,k2_tarski(X83,X82)))|~m1_subset_1(X85,u1_struct_0(X84))|~epred2_3(X84,X83,X82))&(~m1_subset_1(X86,u1_struct_0(X84))|(~r1_orders_2(X84,X86,X83)|~r1_orders_2(X84,X86,X82)|r1_orders_2(X84,X86,X85))|(X85!=k11_lattice3(X84,X83,X82)|~r2_yellow_0(X84,k2_tarski(X83,X82)))|~m1_subset_1(X85,u1_struct_0(X84))|~epred2_3(X84,X83,X82))))&(((X85=k11_lattice3(X84,X83,X82)|(m1_subset_1(esk14_4(X82,X83,X84,X85),u1_struct_0(X84))|~r1_orders_2(X84,X85,X82)|~r1_orders_2(X84,X85,X83))|~m1_subset_1(X85,u1_struct_0(X84))|~epred2_3(X84,X83,X82))&(r2_yellow_0(X84,k2_tarski(X83,X82))|(m1_subset_1(esk14_4(X82,X83,X84,X85),u1_struct_0(X84))|~r1_orders_2(X84,X85,X82)|~r1_orders_2(X84,X85,X83))|~m1_subset_1(X85,u1_struct_0(X84))|~epred2_3(X84,X83,X82)))&((((X85=k11_lattice3(X84,X83,X82)|(r1_orders_2(X84,esk14_4(X82,X83,X84,X85),X83)|~r1_orders_2(X84,X85,X82)|~r1_orders_2(X84,X85,X83))|~m1_subset_1(X85,u1_struct_0(X84))|~epred2_3(X84,X83,X82))&(r2_yellow_0(X84,k2_tarski(X83,X82))|(r1_orders_2(X84,esk14_4(X82,X83,X84,X85),X83)|~r1_orders_2(X84,X85,X82)|~r1_orders_2(X84,X85,X83))|~m1_subset_1(X85,u1_struct_0(X84))|~epred2_3(X84,X83,X82)))&((X85=k11_lattice3(X84,X83,X82)|(r1_orders_2(X84,esk14_4(X82,X83,X84,X85),X82)|~r1_orders_2(X84,X85,X82)|~r1_orders_2(X84,X85,X83))|~m1_subset_1(X85,u1_struct_0(X84))|~epred2_3(X84,X83,X82))&(r2_yellow_0(X84,k2_tarski(X83,X82))|(r1_orders_2(X84,esk14_4(X82,X83,X84,X85),X82)|~r1_orders_2(X84,X85,X82)|~r1_orders_2(X84,X85,X83))|~m1_subset_1(X85,u1_struct_0(X84))|~epred2_3(X84,X83,X82))))&((X85=k11_lattice3(X84,X83,X82)|(~r1_orders_2(X84,esk14_4(X82,X83,X84,X85),X85)|~r1_orders_2(X84,X85,X82)|~r1_orders_2(X84,X85,X83))|~m1_subset_1(X85,u1_struct_0(X84))|~epred2_3(X84,X83,X82))&(r2_yellow_0(X84,k2_tarski(X83,X82))|(~r1_orders_2(X84,esk14_4(X82,X83,X84,X85),X85)|~r1_orders_2(X84,X85,X82)|~r1_orders_2(X84,X85,X83))|~m1_subset_1(X85,u1_struct_0(X84))|~epred2_3(X84,X83,X82)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_140])])])])])).
cnf(c_0_145, negated_conjecture, (~l1_struct_0(esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~r1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_141]), c_0_46])).
cnf(c_0_146, plain, (r1_yellow_0(X1,X2)|~r2_hidden(X3,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_142, c_0_143])).
cnf(c_0_147, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|~r1_orders_2(X1,esk14_4(X3,X2,X1,X4),X4)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~epred2_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_144])).
cnf(c_0_148, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|r1_orders_2(X1,esk14_4(X3,X2,X1,X4),X3)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~epred2_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_144])).
cnf(c_0_149, negated_conjecture, (~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~r1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_57]), c_0_41])])).
cnf(c_0_150, negated_conjecture, (r1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_83]), c_0_113]), c_0_40]), c_0_41]), c_0_65])])).
cnf(c_0_151, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|~epred2_3(X1,X2,X3)|~r1_orders_2(X1,X3,X3)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_147, c_0_148])).
fof(c_0_152, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>epred2_3(X1,X2,X3)))), inference(apply_def,[status(thm)],[t19_yellow_0, c_0_135])).
cnf(c_0_153, negated_conjecture, (~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_149, c_0_150])])).
cnf(c_0_154, negated_conjecture, (r2_yellow_0(esk1_0,k2_tarski(X1,esk2_0))|~epred2_3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_87]), c_0_40])])).
fof(c_0_155, plain, ![X72, X73, X74]:(~v5_orders_2(X72)|~l1_orders_2(X72)|(~m1_subset_1(X73,u1_struct_0(X72))|(~m1_subset_1(X74,u1_struct_0(X72))|epred2_3(X72,X73,X74)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_152])])])).
cnf(c_0_156, negated_conjecture, (~epred2_3(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_154]), c_0_87])])).
cnf(c_0_157, plain, (epred2_3(X1,X2,X3)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_155])).
cnf(c_0_158, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_157]), c_0_40]), c_0_41]), c_0_65])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 159
# Proof object clause steps            : 114
# Proof object formula steps           : 45
# Proof object conjectures             : 58
# Proof object clause conjectures      : 55
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 39
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 63
# Proof object simplifying inferences  : 109
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 91
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 90
# Processed clauses                    : 2220
# ...of these trivial                  : 8
# ...subsumed                          : 1190
# ...remaining for further processing  : 1022
# Other redundant clauses eliminated   : 366
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 172
# Backward-rewritten                   : 46
# Generated clauses                    : 25066
# ...of the previous two non-trivial   : 22870
# Contextual simplify-reflections      : 109
# Paramodulations                      : 24643
# Factorizations                       : 60
# Equation resolutions                 : 366
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 700
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 686
# Current number of unprocessed clauses: 20470
# ...number of literals in the above   : 181170
# Current number of archived formulas  : 0
# Current number of archived clauses   : 306
# Clause-clause subsumption calls (NU) : 114392
# Rec. Clause-clause subsumption calls : 10595
# Non-unit clause-clause subsumptions  : 1423
# Unit Clause-clause subsumption calls : 1727
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 780009

# -------------------------------------------------
# User time                : 0.573 s
# System time              : 0.003 s
# Total time               : 0.576 s
# Maximum resident set size: 28708 pages
