# No SInE strategy applied
# Auto-Mode selected heuristic H_____047_B31_F1_PI_AE_R4_CS_SP_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/34899ea65e2441c89d42b998d9911d5b.E_input', t7_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/34899ea65e2441c89d42b998d9911d5b.E_input', commutativity_k2_xboole_0)).
fof(t11_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k2_xboole_0(X1,X2),X3)=>r1_tarski(X1,X3)), file('data/MPTP/proofs/34899ea65e2441c89d42b998d9911d5b.E_input', t11_xboole_1)).
fof(t8_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X3,X2))=>r1_tarski(k2_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/34899ea65e2441c89d42b998d9911d5b.E_input', t8_xboole_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/34899ea65e2441c89d42b998d9911d5b.E_input', t12_xboole_1)).
fof(t143_zfmisc_1, conjecture, ![X1, X2, X3, X4, X5, X6]:((r1_tarski(X1,k2_zfmisc_1(X3,X4))&r1_tarski(X2,k2_zfmisc_1(X5,X6)))=>r1_tarski(k2_xboole_0(X1,X2),k2_zfmisc_1(k2_xboole_0(X3,X5),k2_xboole_0(X4,X6)))), file('data/MPTP/proofs/34899ea65e2441c89d42b998d9911d5b.E_input', t143_zfmisc_1)).
fof(t120_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&k2_zfmisc_1(X3,k2_xboole_0(X1,X2))=k2_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))), file('data/MPTP/proofs/34899ea65e2441c89d42b998d9911d5b.E_input', t120_zfmisc_1)).
fof(c_0_7, plain, ![X7, X8]:r1_tarski(X7,k2_xboole_0(X7,X8)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_8, plain, ![X9, X10]:k2_xboole_0(X9,X10)=k2_xboole_0(X10,X9), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_9, plain, ![X11, X12, X13]:(~r1_tarski(k2_xboole_0(X11,X12),X13)|r1_tarski(X11,X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_xboole_1])])).
fof(c_0_10, plain, ![X14, X15, X16]:(~r1_tarski(X14,X15)|~r1_tarski(X16,X15)|r1_tarski(k2_xboole_0(X14,X16),X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_xboole_1])])).
cnf(c_0_11, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (r1_tarski(X1,X3)|~r1_tarski(k2_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (r1_tarski(k2_xboole_0(X1,X3),X2)|~r1_tarski(X1,X2)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_11, c_0_12])).
fof(c_0_16, plain, ![X17, X18]:(~r1_tarski(X17,X18)|k2_xboole_0(X17,X18)=X18), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_17, plain, (r1_tarski(X1,X2)|~r1_tarski(k2_xboole_0(X3,X1),X2)), inference(spm,[status(thm)],[c_0_13, c_0_12])).
cnf(c_0_18, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X2,X3))|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_14, c_0_11])).
cnf(c_0_19, plain, (r1_tarski(X1,k2_xboole_0(X2,k2_xboole_0(X1,X3)))), inference(spm,[status(thm)],[c_0_13, c_0_15])).
cnf(c_0_20, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, plain, (r1_tarski(X1,k2_xboole_0(X2,k2_xboole_0(X3,X1)))), inference(spm,[status(thm)],[c_0_17, c_0_15])).
fof(c_0_22, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6]:((r1_tarski(X1,k2_zfmisc_1(X3,X4))&r1_tarski(X2,k2_zfmisc_1(X5,X6)))=>r1_tarski(k2_xboole_0(X1,X2),k2_zfmisc_1(k2_xboole_0(X3,X5),k2_xboole_0(X4,X6))))), inference(assume_negation,[status(cth)],[t143_zfmisc_1])).
cnf(c_0_23, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X2,k2_xboole_0(X1,X3)))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_24, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X2))|~r1_tarski(X1,k2_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_25, plain, (r1_tarski(X1,k2_xboole_0(k2_xboole_0(X2,X1),X3))), inference(spm,[status(thm)],[c_0_17, c_0_11])).
cnf(c_0_26, plain, (k2_xboole_0(X1,k2_xboole_0(X2,k2_xboole_0(X3,X1)))=k2_xboole_0(X2,k2_xboole_0(X3,X1))), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27, plain, (r1_tarski(X1,k2_xboole_0(k2_xboole_0(X1,X2),X3))), inference(spm,[status(thm)],[c_0_13, c_0_11])).
fof(c_0_28, negated_conjecture, ((r1_tarski(esk1_0,k2_zfmisc_1(esk3_0,esk4_0))&r1_tarski(esk2_0,k2_zfmisc_1(esk5_0,esk6_0)))&~r1_tarski(k2_xboole_0(esk1_0,esk2_0),k2_zfmisc_1(k2_xboole_0(esk3_0,esk5_0),k2_xboole_0(esk4_0,esk6_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])).
fof(c_0_29, plain, ![X25, X26, X27]:(k2_zfmisc_1(k2_xboole_0(X25,X26),X27)=k2_xboole_0(k2_zfmisc_1(X25,X27),k2_zfmisc_1(X26,X27))&k2_zfmisc_1(X27,k2_xboole_0(X25,X26))=k2_xboole_0(k2_zfmisc_1(X27,X25),k2_zfmisc_1(X27,X26))), inference(variable_rename,[status(thm)],[t120_zfmisc_1])).
cnf(c_0_30, plain, (k2_xboole_0(k2_xboole_0(X1,X2),k2_xboole_0(X2,k2_xboole_0(X1,X3)))=k2_xboole_0(X2,k2_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_20, c_0_23])).
cnf(c_0_31, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(k2_xboole_0(X3,X1),X2))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_32, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X1,k2_xboole_0(X3,X2)))), inference(spm,[status(thm)],[c_0_23, c_0_26])).
cnf(c_0_33, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(k2_xboole_0(X1,X3),X2))), inference(spm,[status(thm)],[c_0_24, c_0_27])).
cnf(c_0_34, negated_conjecture, (~r1_tarski(k2_xboole_0(esk1_0,esk2_0),k2_zfmisc_1(k2_xboole_0(esk3_0,esk5_0),k2_xboole_0(esk4_0,esk6_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, plain, (k2_zfmisc_1(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, plain, (k2_zfmisc_1(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, plain, (k2_xboole_0(k2_xboole_0(X1,k2_xboole_0(X2,X3)),k2_xboole_0(X2,X1))=k2_xboole_0(X1,k2_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_12, c_0_30])).
cnf(c_0_38, plain, (k2_xboole_0(k2_xboole_0(X1,X2),k2_xboole_0(k2_xboole_0(X3,X1),X2))=k2_xboole_0(k2_xboole_0(X3,X1),X2)), inference(spm,[status(thm)],[c_0_20, c_0_31])).
cnf(c_0_39, plain, (k2_xboole_0(k2_xboole_0(X1,X2),k2_xboole_0(X1,k2_xboole_0(X3,X2)))=k2_xboole_0(X1,k2_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_20, c_0_32])).
cnf(c_0_40, plain, (k2_xboole_0(k2_xboole_0(X1,X2),k2_xboole_0(k2_xboole_0(X1,X3),X2))=k2_xboole_0(k2_xboole_0(X1,X3),X2)), inference(spm,[status(thm)],[c_0_20, c_0_33])).
cnf(c_0_41, negated_conjecture, (r1_tarski(esk1_0,k2_zfmisc_1(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_42, negated_conjecture, (r1_tarski(esk2_0,k2_zfmisc_1(esk5_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_43, negated_conjecture, (~r1_tarski(k2_xboole_0(esk1_0,esk2_0),k2_xboole_0(k2_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk5_0,esk4_0)),k2_xboole_0(k2_zfmisc_1(esk3_0,esk6_0),k2_zfmisc_1(esk5_0,esk6_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_36])).
cnf(c_0_44, plain, (k2_xboole_0(k2_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)),k2_xboole_0(k2_zfmisc_1(X4,X2),k2_zfmisc_1(X4,X3)))=k2_xboole_0(k2_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X4,X2)),k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X4,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_35]), c_0_35]), c_0_36]), c_0_36])).
cnf(c_0_45, plain, (k2_xboole_0(k2_xboole_0(X1,X2),k2_xboole_0(X2,X3))=k2_xboole_0(k2_xboole_0(X1,X2),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
cnf(c_0_46, plain, (k2_xboole_0(k2_xboole_0(X1,X2),k2_xboole_0(X1,X3))=k2_xboole_0(k2_xboole_0(X1,X2),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_40]), c_0_39])).
cnf(c_0_47, negated_conjecture, (k2_xboole_0(esk1_0,k2_zfmisc_1(esk3_0,esk4_0))=k2_zfmisc_1(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_20, c_0_41])).
cnf(c_0_48, negated_conjecture, (k2_xboole_0(esk2_0,k2_zfmisc_1(esk5_0,esk6_0))=k2_zfmisc_1(esk5_0,esk6_0)), inference(spm,[status(thm)],[c_0_20, c_0_42])).
cnf(c_0_49, negated_conjecture, (~r1_tarski(k2_xboole_0(esk1_0,esk2_0),k2_xboole_0(k2_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk3_0,esk6_0)),k2_xboole_0(k2_zfmisc_1(esk5_0,esk4_0),k2_zfmisc_1(esk5_0,esk6_0))))), inference(rw,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_50, plain, (k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X2,k2_xboole_0(X1,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_45]), c_0_46])).
cnf(c_0_51, plain, (k2_xboole_0(X1,k2_xboole_0(X2,k2_xboole_0(X1,X3)))=k2_xboole_0(X2,k2_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_20, c_0_19])).
cnf(c_0_52, negated_conjecture, (r1_tarski(esk1_0,k2_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),X1))), inference(spm,[status(thm)],[c_0_27, c_0_47])).
cnf(c_0_53, negated_conjecture, (r1_tarski(esk2_0,X1)|~r1_tarski(k2_zfmisc_1(esk5_0,esk6_0),X1)), inference(spm,[status(thm)],[c_0_13, c_0_48])).
cnf(c_0_54, plain, (r1_tarski(X1,k2_xboole_0(X2,k2_xboole_0(X3,k2_xboole_0(X4,X1))))), inference(spm,[status(thm)],[c_0_17, c_0_21])).
cnf(c_0_55, negated_conjecture, (~r1_tarski(k2_xboole_0(esk1_0,esk2_0),k2_xboole_0(k2_zfmisc_1(esk3_0,esk6_0),k2_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_xboole_0(k2_zfmisc_1(esk5_0,esk4_0),k2_zfmisc_1(esk5_0,esk6_0)))))), inference(rw,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_56, plain, (k2_xboole_0(X1,k2_xboole_0(X2,X3))=k2_xboole_0(X2,k2_xboole_0(X1,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_50]), c_0_50]), c_0_51])).
cnf(c_0_57, negated_conjecture, (r1_tarski(k2_xboole_0(X1,esk1_0),k2_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),X2))|~r1_tarski(X1,k2_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),X2))), inference(spm,[status(thm)],[c_0_14, c_0_52])).
cnf(c_0_58, negated_conjecture, (r1_tarski(esk2_0,k2_xboole_0(X1,k2_xboole_0(X2,k2_xboole_0(X3,k2_zfmisc_1(esk5_0,esk6_0)))))), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_59, negated_conjecture, (~r1_tarski(k2_xboole_0(esk1_0,esk2_0),k2_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_xboole_0(k2_zfmisc_1(esk3_0,esk6_0),k2_xboole_0(k2_zfmisc_1(esk5_0,esk4_0),k2_zfmisc_1(esk5_0,esk6_0)))))), inference(rw,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_60, negated_conjecture, (r1_tarski(k2_xboole_0(esk1_0,esk2_0),k2_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_xboole_0(X1,k2_xboole_0(X2,k2_zfmisc_1(esk5_0,esk6_0)))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_12])).
cnf(c_0_61, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_60])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 62
# Proof object clause steps            : 47
# Proof object formula steps           : 15
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 30
# Proof object simplifying inferences  : 20
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 10
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 10
# Processed clauses                    : 10845
# ...of these trivial                  : 4507
# ...subsumed                          : 5620
# ...remaining for further processing  : 718
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 110
# Generated clauses                    : 152330
# ...of the previous two non-trivial   : 88090
# Contextual simplify-reflections      : 0
# Paramodulations                      : 152330
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 608
#    Positive orientable unit clauses  : 410
#    Positive unorientable unit clauses: 17
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 181
# Current number of unprocessed clauses: 70506
# ...number of literals in the above   : 87099
# Current number of archived formulas  : 0
# Current number of archived clauses   : 110
# Clause-clause subsumption calls (NU) : 56948
# Rec. Clause-clause subsumption calls : 56948
# Non-unit clause-clause subsumptions  : 4273
# Unit Clause-clause subsumption calls : 487
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 13820
# BW rewrite match successes           : 3630
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2066335

# -------------------------------------------------
# User time                : 0.898 s
# System time              : 0.023 s
# Total time               : 0.921 s
# Maximum resident set size: 64892 pages
