# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d3_xboole_0, axiom, ![X1, X2, X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2)))), file('data/MPTP/proofs/3555f766284f423e80214ffbe2acd875.E_input', d3_xboole_0)).
fof(t127_zfmisc_1, conjecture, ![X1, X2, X3, X4]:((r1_xboole_0(X1,X2)|r1_xboole_0(X3,X4))=>r1_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), file('data/MPTP/proofs/3555f766284f423e80214ffbe2acd875.E_input', t127_zfmisc_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/3555f766284f423e80214ffbe2acd875.E_input', d3_tarski)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/3555f766284f423e80214ffbe2acd875.E_input', d7_xboole_0)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/3555f766284f423e80214ffbe2acd875.E_input', t48_xboole_1)).
fof(t122_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k3_xboole_0(X1,X2),X3)=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&k2_zfmisc_1(X3,k3_xboole_0(X1,X2))=k3_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))), file('data/MPTP/proofs/3555f766284f423e80214ffbe2acd875.E_input', t122_zfmisc_1)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/3555f766284f423e80214ffbe2acd875.E_input', t83_xboole_1)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/3555f766284f423e80214ffbe2acd875.E_input', t2_boole)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/3555f766284f423e80214ffbe2acd875.E_input', t3_boole)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/3555f766284f423e80214ffbe2acd875.E_input', t113_zfmisc_1)).
fof(t63_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_xboole_0(X2,X3))=>r1_xboole_0(X1,X3)), file('data/MPTP/proofs/3555f766284f423e80214ffbe2acd875.E_input', t63_xboole_1)).
fof(t70_xboole_1, axiom, ![X1, X2, X3]:(~((~(r1_xboole_0(X1,k2_xboole_0(X2,X3)))&(r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3))))&~((~((r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3)))&r1_xboole_0(X1,k2_xboole_0(X2,X3))))), file('data/MPTP/proofs/3555f766284f423e80214ffbe2acd875.E_input', t70_xboole_1)).
fof(t120_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&k2_zfmisc_1(X3,k2_xboole_0(X1,X2))=k2_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))), file('data/MPTP/proofs/3555f766284f423e80214ffbe2acd875.E_input', t120_zfmisc_1)).
fof(c_0_13, plain, ![X5, X6, X7, X8, X9, X10, X11, X12]:(((~r2_hidden(X8,X7)|(r2_hidden(X8,X5)|r2_hidden(X8,X6))|X7!=k2_xboole_0(X5,X6))&((~r2_hidden(X9,X5)|r2_hidden(X9,X7)|X7!=k2_xboole_0(X5,X6))&(~r2_hidden(X9,X6)|r2_hidden(X9,X7)|X7!=k2_xboole_0(X5,X6))))&(((~r2_hidden(esk1_3(X10,X11,X12),X10)|~r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k2_xboole_0(X10,X11))&(~r2_hidden(esk1_3(X10,X11,X12),X11)|~r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k2_xboole_0(X10,X11)))&(r2_hidden(esk1_3(X10,X11,X12),X12)|(r2_hidden(esk1_3(X10,X11,X12),X10)|r2_hidden(esk1_3(X10,X11,X12),X11))|X12=k2_xboole_0(X10,X11)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])).
fof(c_0_14, negated_conjecture, ~(![X1, X2, X3, X4]:((r1_xboole_0(X1,X2)|r1_xboole_0(X3,X4))=>r1_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)))), inference(assume_negation,[status(cth)],[t127_zfmisc_1])).
fof(c_0_15, plain, ![X18, X19, X20, X21, X22]:((~r1_tarski(X18,X19)|(~r2_hidden(X20,X18)|r2_hidden(X20,X19)))&((r2_hidden(esk6_2(X21,X22),X21)|r1_tarski(X21,X22))&(~r2_hidden(esk6_2(X21,X22),X22)|r1_tarski(X21,X22)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_16, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_17, plain, ![X24, X25]:((~r1_xboole_0(X24,X25)|k3_xboole_0(X24,X25)=k1_xboole_0)&(k3_xboole_0(X24,X25)!=k1_xboole_0|r1_xboole_0(X24,X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
fof(c_0_18, plain, ![X26, X27]:k4_xboole_0(X26,k4_xboole_0(X26,X27))=k3_xboole_0(X26,X27), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_19, plain, ![X28, X29, X30]:(k2_zfmisc_1(k3_xboole_0(X28,X29),X30)=k3_xboole_0(k2_zfmisc_1(X28,X30),k2_zfmisc_1(X29,X30))&k2_zfmisc_1(X30,k3_xboole_0(X28,X29))=k3_xboole_0(k2_zfmisc_1(X30,X28),k2_zfmisc_1(X30,X29))), inference(variable_rename,[status(thm)],[t122_zfmisc_1])).
fof(c_0_20, plain, ![X31, X32]:((~r1_xboole_0(X31,X32)|k4_xboole_0(X31,X32)=X31)&(k4_xboole_0(X31,X32)!=X31|r1_xboole_0(X31,X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
fof(c_0_21, negated_conjecture, ((r1_xboole_0(esk2_0,esk3_0)|r1_xboole_0(esk4_0,esk5_0))&~r1_xboole_0(k2_zfmisc_1(esk2_0,esk4_0),k2_zfmisc_1(esk3_0,esk5_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
fof(c_0_22, plain, ![X33]:k3_xboole_0(X33,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
cnf(c_0_23, plain, (r1_tarski(X1,X2)|~r2_hidden(esk6_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_25, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, plain, (k2_zfmisc_1(X1,k3_xboole_0(X2,X3))=k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, plain, (k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, negated_conjecture, (r1_xboole_0(esk2_0,esk3_0)|r1_xboole_0(esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_31, plain, ![X34]:k4_xboole_0(X34,k1_xboole_0)=X34, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_32, plain, ![X35, X36]:((k2_zfmisc_1(X35,X36)!=k1_xboole_0|(X35=k1_xboole_0|X36=k1_xboole_0))&((X35!=k1_xboole_0|k2_zfmisc_1(X35,X36)=k1_xboole_0)&(X36!=k1_xboole_0|k2_zfmisc_1(X35,X36)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_33, plain, ![X37, X38, X39]:(~r1_tarski(X37,X38)|~r1_xboole_0(X38,X39)|r1_xboole_0(X37,X39)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_xboole_1])])).
cnf(c_0_34, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r2_hidden(esk6_2(X1,k2_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_35, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_36, plain, ![X40, X41, X42]:((r1_xboole_0(X40,k2_xboole_0(X41,X42))|(~r1_xboole_0(X40,X41)|~r1_xboole_0(X40,X42)))&((r1_xboole_0(X40,X41)|~r1_xboole_0(X40,k2_xboole_0(X41,X42)))&(r1_xboole_0(X40,X42)|~r1_xboole_0(X40,k2_xboole_0(X41,X42))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t70_xboole_1])])])])).
fof(c_0_37, plain, ![X43, X44, X45]:(k2_zfmisc_1(k2_xboole_0(X43,X44),X45)=k2_xboole_0(k2_zfmisc_1(X43,X45),k2_zfmisc_1(X44,X45))&k2_zfmisc_1(X45,k2_xboole_0(X43,X44))=k2_xboole_0(k2_zfmisc_1(X45,X43),k2_zfmisc_1(X45,X44))), inference(variable_rename,[status(thm)],[t120_zfmisc_1])).
cnf(c_0_38, plain, (r1_xboole_0(X1,X2)|k4_xboole_0(X1,k4_xboole_0(X1,X2))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_39, plain, (k2_zfmisc_1(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))=k4_xboole_0(k2_zfmisc_1(X1,X2),k4_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_26]), c_0_26])).
cnf(c_0_40, negated_conjecture, (k4_xboole_0(esk4_0,esk5_0)=esk4_0|r1_xboole_0(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_41, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k1_xboole_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_30, c_0_26])).
cnf(c_0_42, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_43, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_44, plain, (r1_xboole_0(X1,X3)|~r1_tarski(X1,X2)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_45, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_46, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_47, plain, (k2_zfmisc_1(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_48, plain, (r1_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))|k2_zfmisc_1(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_49, negated_conjecture, (k4_xboole_0(esk4_0,esk5_0)=esk4_0|k4_xboole_0(esk2_0,esk3_0)=esk2_0), inference(spm,[status(thm)],[c_0_28, c_0_40])).
cnf(c_0_50, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_51, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_43])).
cnf(c_0_52, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(k2_xboole_0(X3,X1),X2)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_53, plain, (r1_xboole_0(X1,k2_zfmisc_1(X2,X3))|~r1_xboole_0(X1,k2_zfmisc_1(k2_xboole_0(X2,X4),X3))), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_54, negated_conjecture, (k4_xboole_0(esk2_0,esk3_0)=esk2_0|r1_xboole_0(k2_zfmisc_1(X1,esk4_0),k2_zfmisc_1(X1,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50]), c_0_51])])).
cnf(c_0_55, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),X3)=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_56, plain, (r1_xboole_0(k2_zfmisc_1(X1,X2),X3)|~r1_xboole_0(k2_zfmisc_1(k2_xboole_0(X4,X1),X2),X3)), inference(spm,[status(thm)],[c_0_52, c_0_47])).
cnf(c_0_57, negated_conjecture, (k4_xboole_0(esk2_0,esk3_0)=esk2_0|r1_xboole_0(k2_zfmisc_1(k2_xboole_0(X1,X2),esk4_0),k2_zfmisc_1(X1,esk5_0))), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_58, plain, (k2_zfmisc_1(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X3)=k4_xboole_0(k2_zfmisc_1(X1,X3),k4_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_26]), c_0_26])).
cnf(c_0_59, negated_conjecture, (~r1_xboole_0(k2_zfmisc_1(esk2_0,esk4_0),k2_zfmisc_1(esk3_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_60, negated_conjecture, (k4_xboole_0(esk2_0,esk3_0)=esk2_0|r1_xboole_0(k2_zfmisc_1(X1,esk4_0),k2_zfmisc_1(X2,esk5_0))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_61, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_62, plain, (k2_zfmisc_1(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_63, plain, (r1_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))|k2_zfmisc_1(k4_xboole_0(X1,k4_xboole_0(X1,X3)),X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_38, c_0_58])).
cnf(c_0_64, negated_conjecture, (k4_xboole_0(esk2_0,esk3_0)=esk2_0), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_65, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_61])).
cnf(c_0_66, plain, (r1_xboole_0(X1,k2_zfmisc_1(X2,X3))|~r1_xboole_0(X1,k2_zfmisc_1(X2,k2_xboole_0(X3,X4)))), inference(spm,[status(thm)],[c_0_46, c_0_62])).
cnf(c_0_67, negated_conjecture, (r1_xboole_0(k2_zfmisc_1(esk2_0,X1),k2_zfmisc_1(esk3_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_50]), c_0_65])])).
cnf(c_0_68, plain, (r1_xboole_0(k2_zfmisc_1(X1,X2),X3)|~r1_xboole_0(k2_zfmisc_1(X1,k2_xboole_0(X4,X2)),X3)), inference(spm,[status(thm)],[c_0_52, c_0_62])).
cnf(c_0_69, negated_conjecture, (r1_xboole_0(k2_zfmisc_1(esk2_0,k2_xboole_0(X1,X2)),k2_zfmisc_1(esk3_0,X1))), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_70, negated_conjecture, (r1_xboole_0(k2_zfmisc_1(esk2_0,X1),k2_zfmisc_1(esk3_0,X2))), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_71, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_70])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 72
# Proof object clause steps            : 45
# Proof object formula steps           : 27
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 28
# Processed clauses                    : 1349
# ...of these trivial                  : 22
# ...subsumed                          : 801
# ...remaining for further processing  : 526
# Other redundant clauses eliminated   : 37
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 25
# Backward-rewritten                   : 124
# Generated clauses                    : 14083
# ...of the previous two non-trivial   : 10146
# Contextual simplify-reflections      : 1
# Paramodulations                      : 13993
# Factorizations                       : 53
# Equation resolutions                 : 37
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 344
#    Positive orientable unit clauses  : 65
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 277
# Current number of unprocessed clauses: 8797
# ...number of literals in the above   : 23097
# Current number of archived formulas  : 0
# Current number of archived clauses   : 178
# Clause-clause subsumption calls (NU) : 13758
# Rec. Clause-clause subsumption calls : 11926
# Non-unit clause-clause subsumptions  : 818
# Unit Clause-clause subsumption calls : 665
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 90
# BW rewrite match successes           : 21
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 193017

# -------------------------------------------------
# User time                : 0.101 s
# System time              : 0.007 s
# Total time               : 0.108 s
# Maximum resident set size: 28472 pages
