# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t49_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v15_waybel_0(X2,X1)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&X2=k6_waybel_0(X1,X3))))), file('data/MPTP/proofs/355d53feb06d40108e3c8aa1b6d18825.E_input', t49_waybel_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/355d53feb06d40108e3c8aa1b6d18825.E_input', t4_subset)).
fof(d22_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v15_waybel_0(X2,X1)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_hidden(X3,X2)&r1_lattice3(X1,X2,X3)))))), file('data/MPTP/proofs/355d53feb06d40108e3c8aa1b6d18825.E_input', d22_waybel_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/355d53feb06d40108e3c8aa1b6d18825.E_input', d3_tarski)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/355d53feb06d40108e3c8aa1b6d18825.E_input', d8_lattice3)).
fof(dt_k6_waybel_0, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k6_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/355d53feb06d40108e3c8aa1b6d18825.E_input', dt_k6_waybel_0)).
fof(d20_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v13_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r1_orders_2(X1,X3,X4))=>r2_hidden(X4,X2))))))), file('data/MPTP/proofs/355d53feb06d40108e3c8aa1b6d18825.E_input', d20_waybel_0)).
fof(t18_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k6_waybel_0(X1,X2))<=>r1_orders_2(X1,X2,X3))))), file('data/MPTP/proofs/355d53feb06d40108e3c8aa1b6d18825.E_input', t18_waybel_0)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/355d53feb06d40108e3c8aa1b6d18825.E_input', d10_xboole_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/355d53feb06d40108e3c8aa1b6d18825.E_input', t24_orders_2)).
fof(c_0_10, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v15_waybel_0(X2,X1)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&X2=k6_waybel_0(X1,X3)))))), inference(assume_negation,[status(cth)],[t49_waybel_0])).
fof(c_0_11, plain, ![X9, X10, X11]:(~r2_hidden(X9,X10)|~m1_subset_1(X10,k1_zfmisc_1(X11))|m1_subset_1(X9,X11)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_12, negated_conjecture, ![X7]:((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0))))&((~v1_xboole_0(esk2_0)&(v2_waybel_0(esk2_0,esk1_0)&(v13_waybel_0(esk2_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))))&((~v15_waybel_0(esk2_0,esk1_0)|(~m1_subset_1(X7,u1_struct_0(esk1_0))|esk2_0!=k6_waybel_0(esk1_0,X7)))&((m1_subset_1(esk3_0,u1_struct_0(esk1_0))|v15_waybel_0(esk2_0,esk1_0))&(esk2_0=k6_waybel_0(esk1_0,esk3_0)|v15_waybel_0(esk2_0,esk1_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_10])])])])])])).
fof(c_0_13, plain, ![X12, X13, X15]:(((m1_subset_1(esk4_2(X12,X13),u1_struct_0(X12))|~v15_waybel_0(X13,X12)|(v1_xboole_0(X13)|(~v2_waybel_0(X13,X12)|(~v13_waybel_0(X13,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12))))))|(v2_struct_0(X12)|(~v3_orders_2(X12)|(~v4_orders_2(X12)|~l1_orders_2(X12)))))&((r2_hidden(esk4_2(X12,X13),X13)|~v15_waybel_0(X13,X12)|(v1_xboole_0(X13)|(~v2_waybel_0(X13,X12)|(~v13_waybel_0(X13,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12))))))|(v2_struct_0(X12)|(~v3_orders_2(X12)|(~v4_orders_2(X12)|~l1_orders_2(X12)))))&(r1_lattice3(X12,X13,esk4_2(X12,X13))|~v15_waybel_0(X13,X12)|(v1_xboole_0(X13)|(~v2_waybel_0(X13,X12)|(~v13_waybel_0(X13,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12))))))|(v2_struct_0(X12)|(~v3_orders_2(X12)|(~v4_orders_2(X12)|~l1_orders_2(X12)))))))&(~m1_subset_1(X15,u1_struct_0(X12))|(~r2_hidden(X15,X13)|~r1_lattice3(X12,X13,X15))|v15_waybel_0(X13,X12)|(v1_xboole_0(X13)|(~v2_waybel_0(X13,X12)|(~v13_waybel_0(X13,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12))))))|(v2_struct_0(X12)|(~v3_orders_2(X12)|(~v4_orders_2(X12)|~l1_orders_2(X12)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d22_waybel_0])])])])])])).
cnf(c_0_14, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_16, plain, ![X16, X17, X18, X19, X20]:((~r1_tarski(X16,X17)|(~r2_hidden(X18,X16)|r2_hidden(X18,X17)))&((r2_hidden(esk5_2(X19,X20),X19)|r1_tarski(X19,X20))&(~r2_hidden(esk5_2(X19,X20),X20)|r1_tarski(X19,X20)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_17, plain, (m1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~v15_waybel_0(X2,X1)|~v2_waybel_0(X2,X1)|~v13_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, negated_conjecture, (v13_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, negated_conjecture, (v2_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_24, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_25, plain, ![X22, X23, X24, X25]:((~r1_lattice3(X22,X23,X24)|(~m1_subset_1(X25,u1_struct_0(X22))|(~r2_hidden(X25,X23)|r1_orders_2(X22,X24,X25)))|~m1_subset_1(X24,u1_struct_0(X22))|~l1_orders_2(X22))&((m1_subset_1(esk6_3(X22,X23,X24),u1_struct_0(X22))|r1_lattice3(X22,X23,X24)|~m1_subset_1(X24,u1_struct_0(X22))|~l1_orders_2(X22))&((r2_hidden(esk6_3(X22,X23,X24),X23)|r1_lattice3(X22,X23,X24)|~m1_subset_1(X24,u1_struct_0(X22))|~l1_orders_2(X22))&(~r1_orders_2(X22,X24,esk6_3(X22,X23,X24))|r1_lattice3(X22,X23,X24)|~m1_subset_1(X24,u1_struct_0(X22))|~l1_orders_2(X22))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
cnf(c_0_26, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_27, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_28, plain, ![X27, X28]:(v2_struct_0(X27)|~l1_orders_2(X27)|~m1_subset_1(X28,u1_struct_0(X27))|m1_subset_1(k6_waybel_0(X27,X28),k1_zfmisc_1(u1_struct_0(X27)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_waybel_0])])])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk4_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~v15_waybel_0(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_15]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22])]), c_0_23]), c_0_24])).
cnf(c_0_30, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))|v15_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_31, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, negated_conjecture, (r1_tarski(esk2_0,X1)|m1_subset_1(esk5_2(esk2_0,X1),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_33, plain, (r1_lattice3(X1,X2,esk4_2(X1,X2))|v1_xboole_0(X2)|v2_struct_0(X1)|~v15_waybel_0(X2,X1)|~v2_waybel_0(X2,X1)|~v13_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_34, plain, ![X29, X30, X31, X32]:((~v13_waybel_0(X30,X29)|(~m1_subset_1(X31,u1_struct_0(X29))|(~m1_subset_1(X32,u1_struct_0(X29))|(~r2_hidden(X31,X30)|~r1_orders_2(X29,X31,X32)|r2_hidden(X32,X30))))|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))|~l1_orders_2(X29))&((m1_subset_1(esk7_2(X29,X30),u1_struct_0(X29))|v13_waybel_0(X30,X29)|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))|~l1_orders_2(X29))&((m1_subset_1(esk8_2(X29,X30),u1_struct_0(X29))|v13_waybel_0(X30,X29)|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))|~l1_orders_2(X29))&(((r2_hidden(esk7_2(X29,X30),X30)|v13_waybel_0(X30,X29)|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))|~l1_orders_2(X29))&(r1_orders_2(X29,esk7_2(X29,X30),esk8_2(X29,X30))|v13_waybel_0(X30,X29)|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))|~l1_orders_2(X29)))&(~r2_hidden(esk8_2(X29,X30),X30)|v13_waybel_0(X30,X29)|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))|~l1_orders_2(X29)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d20_waybel_0])])])])])).
cnf(c_0_35, plain, (v2_struct_0(X1)|m1_subset_1(k6_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36, negated_conjecture, (m1_subset_1(esk4_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
fof(c_0_37, plain, ![X35, X36, X37]:((~r2_hidden(X37,k6_waybel_0(X35,X36))|r1_orders_2(X35,X36,X37)|~m1_subset_1(X37,u1_struct_0(X35))|~m1_subset_1(X36,u1_struct_0(X35))|(v2_struct_0(X35)|~l1_orders_2(X35)))&(~r1_orders_2(X35,X36,X37)|r2_hidden(X37,k6_waybel_0(X35,X36))|~m1_subset_1(X37,u1_struct_0(X35))|~m1_subset_1(X36,u1_struct_0(X35))|(v2_struct_0(X35)|~l1_orders_2(X35)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t18_waybel_0])])])])])).
cnf(c_0_38, negated_conjecture, (r1_orders_2(esk1_0,X1,esk5_2(esk2_0,X2))|r1_tarski(esk2_0,X2)|~r1_lattice3(esk1_0,X3,X1)|~r2_hidden(esk5_2(esk2_0,X2),X3)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_20])])).
cnf(c_0_39, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0))|~v15_waybel_0(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_15]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22])]), c_0_23]), c_0_24])).
cnf(c_0_40, plain, (r2_hidden(X4,X1)|~v13_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,X1)|~r1_orders_2(X2,X3,X4)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_41, negated_conjecture, (m1_subset_1(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_20])]), c_0_24])).
cnf(c_0_42, plain, (r2_hidden(X3,k6_waybel_0(X1,X2))|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_43, negated_conjecture, (r1_orders_2(esk1_0,X1,esk5_2(esk2_0,X2))|r1_tarski(esk2_0,X2)|~r1_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_38, c_0_27])).
cnf(c_0_44, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0))|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_39, c_0_30])).
cnf(c_0_45, plain, (r2_hidden(X1,X2)|~r1_orders_2(X3,X4,X1)|~r2_hidden(X4,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~v13_waybel_0(X2,X3)|~l1_orders_2(X3)), inference(csr,[status(thm)],[c_0_40, c_0_14])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))|m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_14, c_0_41])).
cnf(c_0_47, plain, (r1_orders_2(X2,X3,X1)|v2_struct_0(X2)|~r2_hidden(X1,k6_waybel_0(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_48, plain, (r2_hidden(esk4_2(X1,X2),X2)|v1_xboole_0(X2)|v2_struct_0(X1)|~v15_waybel_0(X2,X1)|~v2_waybel_0(X2,X1)|~v13_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_49, negated_conjecture, (r1_tarski(esk2_0,X1)|r2_hidden(esk5_2(esk2_0,X1),k6_waybel_0(esk1_0,X2))|~r1_orders_2(esk1_0,X2,esk5_2(esk2_0,X1))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_32]), c_0_20])]), c_0_24])).
cnf(c_0_50, negated_conjecture, (r1_orders_2(esk1_0,esk4_2(esk1_0,esk2_0),esk5_2(esk2_0,X1))|r1_tarski(esk2_0,X1)|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_36])).
cnf(c_0_51, negated_conjecture, (~v15_waybel_0(esk2_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|esk2_0!=k6_waybel_0(esk1_0,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_52, negated_conjecture, (r2_hidden(X1,esk2_0)|~r1_orders_2(esk1_0,X2,X1)|~r2_hidden(X2,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_15]), c_0_18]), c_0_20])])).
cnf(c_0_53, negated_conjecture, (r1_tarski(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1)|m1_subset_1(esk5_2(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1),u1_struct_0(esk1_0))|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_46, c_0_27])).
cnf(c_0_54, plain, (r1_orders_2(X1,X2,esk5_2(k6_waybel_0(X1,X2),X3))|r1_tarski(k6_waybel_0(X1,X2),X3)|v2_struct_0(X1)|~m1_subset_1(esk5_2(k6_waybel_0(X1,X2),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_47, c_0_27])).
cnf(c_0_55, negated_conjecture, (r2_hidden(esk4_2(esk1_0,esk2_0),esk2_0)|~v15_waybel_0(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_15]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22])]), c_0_23]), c_0_24])).
fof(c_0_56, plain, ![X38, X39]:(((r1_tarski(X38,X39)|X38!=X39)&(r1_tarski(X39,X38)|X38!=X39))&(~r1_tarski(X38,X39)|~r1_tarski(X39,X38)|X38=X39)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_57, plain, (r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_58, negated_conjecture, (r1_tarski(esk2_0,X1)|r2_hidden(esk5_2(esk2_0,X1),k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)))|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_36])).
cnf(c_0_59, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))|k6_waybel_0(esk1_0,X1)!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_51, c_0_30])).
cnf(c_0_60, negated_conjecture, (r1_tarski(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1)|r2_hidden(esk5_2(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1),esk2_0)|m1_subset_1(esk3_0,u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,X2,esk5_2(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1))|~r2_hidden(X2,esk2_0)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_61, negated_conjecture, (r1_orders_2(esk1_0,esk4_2(esk1_0,esk2_0),esk5_2(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1))|r1_tarski(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1)|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_53]), c_0_20])]), c_0_24]), c_0_36])).
cnf(c_0_62, negated_conjecture, (r2_hidden(esk4_2(esk1_0,esk2_0),esk2_0)|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_55, c_0_30])).
cnf(c_0_63, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_64, negated_conjecture, (r1_tarski(esk2_0,k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)))|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_65, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))|k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0))!=esk2_0), inference(spm,[status(thm)],[c_0_59, c_0_36])).
cnf(c_0_66, negated_conjecture, (r1_tarski(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1)|r2_hidden(esk5_2(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1),esk2_0)|m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62])).
cnf(c_0_67, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))|~r1_tarski(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_65])).
fof(c_0_68, plain, ![X40, X41]:(v2_struct_0(X40)|(~v3_orders_2(X40)|~l1_orders_2(X40))|(~m1_subset_1(X41,u1_struct_0(X40))|r1_orders_2(X40,X41,X41))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
cnf(c_0_69, negated_conjecture, (esk2_0=k6_waybel_0(esk1_0,esk3_0)|v15_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_70, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_71, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_66]), c_0_67])).
cnf(c_0_72, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_73, plain, (v15_waybel_0(X3,X2)|v1_xboole_0(X3)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r1_lattice3(X2,X3,X1)|~v2_waybel_0(X3,X2)|~v13_waybel_0(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v3_orders_2(X2)|~v4_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_74, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,X1)|v15_waybel_0(esk2_0,esk1_0)|~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_69]), c_0_20])]), c_0_24]), c_0_26]), c_0_30])).
cnf(c_0_75, negated_conjecture, (r1_lattice3(esk1_0,X1,esk3_0)|r2_hidden(esk6_3(esk1_0,X1,esk3_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_20])])).
cnf(c_0_76, negated_conjecture, (r2_hidden(esk3_0,k6_waybel_0(esk1_0,X1))|~r1_orders_2(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_71]), c_0_20])]), c_0_24])).
cnf(c_0_77, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_71]), c_0_20]), c_0_22])]), c_0_24])).
cnf(c_0_78, plain, (v15_waybel_0(X1,X2)|v1_xboole_0(X1)|v2_struct_0(X2)|~r1_lattice3(X2,X1,X3)|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v13_waybel_0(X1,X2)|~v2_waybel_0(X1,X2)|~l1_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[c_0_73, c_0_14])).
cnf(c_0_79, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk6_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_80, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk6_3(esk1_0,esk2_0,esk3_0))|r1_lattice3(esk1_0,esk2_0,esk3_0)|v15_waybel_0(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_81, negated_conjecture, (r2_hidden(esk3_0,k6_waybel_0(esk1_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_71])])).
cnf(c_0_82, negated_conjecture, (v15_waybel_0(esk2_0,esk1_0)|~r1_lattice3(esk1_0,esk2_0,X1)|~r2_hidden(X1,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_15]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22])]), c_0_23]), c_0_24])).
cnf(c_0_83, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk3_0)|v15_waybel_0(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_71]), c_0_20])])).
cnf(c_0_84, negated_conjecture, (r2_hidden(esk3_0,esk2_0)|v15_waybel_0(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_81, c_0_69])).
cnf(c_0_85, negated_conjecture, (v15_waybel_0(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_84])).
cnf(c_0_86, negated_conjecture, (m1_subset_1(esk4_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_85])])).
cnf(c_0_87, negated_conjecture, (m1_subset_1(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_86]), c_0_20])]), c_0_24])).
cnf(c_0_88, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_85])])).
cnf(c_0_89, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_14, c_0_87])).
cnf(c_0_90, negated_conjecture, (r1_orders_2(esk1_0,esk4_2(esk1_0,esk2_0),esk5_2(esk2_0,X1))|r1_tarski(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_88]), c_0_86])])).
cnf(c_0_91, negated_conjecture, (r1_tarski(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1)|m1_subset_1(esk5_2(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_89, c_0_27])).
cnf(c_0_92, negated_conjecture, (r1_tarski(esk2_0,X1)|r2_hidden(esk5_2(esk2_0,X1),k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_90]), c_0_86])])).
cnf(c_0_93, negated_conjecture, (k6_waybel_0(esk1_0,X1)!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_85])])).
cnf(c_0_94, negated_conjecture, (r1_tarski(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1)|r2_hidden(esk5_2(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1),esk2_0)|~r1_orders_2(esk1_0,X2,esk5_2(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1))|~r2_hidden(X2,esk2_0)), inference(spm,[status(thm)],[c_0_52, c_0_91])).
cnf(c_0_95, negated_conjecture, (r1_orders_2(esk1_0,esk4_2(esk1_0,esk2_0),esk5_2(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1))|r1_tarski(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_91]), c_0_86]), c_0_20])]), c_0_24])).
cnf(c_0_96, negated_conjecture, (r2_hidden(esk4_2(esk1_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_85])])).
cnf(c_0_97, negated_conjecture, (r1_tarski(esk2_0,k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_57, c_0_92])).
cnf(c_0_98, negated_conjecture, (k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0))!=esk2_0), inference(spm,[status(thm)],[c_0_93, c_0_86])).
cnf(c_0_99, negated_conjecture, (r1_tarski(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1)|r2_hidden(esk5_2(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),X1),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_96])])).
cnf(c_0_100, negated_conjecture, (~r1_tarski(k6_waybel_0(esk1_0,esk4_2(esk1_0,esk2_0)),esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_97]), c_0_98])).
cnf(c_0_101, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_99]), c_0_100]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 102
# Proof object clause steps            : 81
# Proof object formula steps           : 21
# Proof object conjectures             : 65
# Proof object clause conjectures      : 62
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 48
# Proof object simplifying inferences  : 97
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 36
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 36
# Processed clauses                    : 562
# ...of these trivial                  : 2
# ...subsumed                          : 102
# ...remaining for further processing  : 458
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 58
# Backward-rewritten                   : 78
# Generated clauses                    : 1367
# ...of the previous two non-trivial   : 1324
# Contextual simplify-reflections      : 31
# Paramodulations                      : 1365
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 285
#    Positive orientable unit clauses  : 19
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 261
# Current number of unprocessed clauses: 726
# ...number of literals in the above   : 3570
# Current number of archived formulas  : 0
# Current number of archived clauses   : 171
# Clause-clause subsumption calls (NU) : 19796
# Rec. Clause-clause subsumption calls : 3750
# Non-unit clause-clause subsumptions  : 185
# Unit Clause-clause subsumption calls : 457
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 19
# BW rewrite match successes           : 7
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 56079

# -------------------------------------------------
# User time                : 0.042 s
# System time              : 0.000 s
# Total time               : 0.042 s
# Maximum resident set size: 28832 pages
