# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/35d128e821994ca6bf5ed5300c5d736a.E_input', l3_subset_1)).
fof(dt_k2_pre_topc, axiom, ![X1, X2]:((l1_pre_topc(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k2_pre_topc(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/35d128e821994ca6bf5ed5300c5d736a.E_input', dt_k2_pre_topc)).
fof(t45_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(r2_hidden(X3,u1_struct_0(X1))=>(r2_hidden(X3,k2_pre_topc(X1,X2))<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>((v4_pre_topc(X4,X1)&r1_tarski(X2,X4))=>r2_hidden(X3,X4))))))), file('data/MPTP/proofs/35d128e821994ca6bf5ed5300c5d736a.E_input', t45_pre_topc)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/35d128e821994ca6bf5ed5300c5d736a.E_input', d3_tarski)).
fof(t44_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X3,X2)=>v4_pre_topc(X3,X1)))=>v4_pre_topc(k6_setfam_1(u1_struct_0(X1),X2),X1)))), file('data/MPTP/proofs/35d128e821994ca6bf5ed5300c5d736a.E_input', t44_pre_topc)).
fof(t46_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>?[X3]:(m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))&(![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X4,X3)<=>(v4_pre_topc(X4,X1)&r1_tarski(X2,X4))))&k2_pre_topc(X1,X2)=k6_setfam_1(u1_struct_0(X1),X3))))), file('data/MPTP/proofs/35d128e821994ca6bf5ed5300c5d736a.E_input', t46_pre_topc)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/35d128e821994ca6bf5ed5300c5d736a.E_input', d10_xboole_0)).
fof(t48_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>r1_tarski(X2,k2_pre_topc(X1,X2)))), file('data/MPTP/proofs/35d128e821994ca6bf5ed5300c5d736a.E_input', t48_pre_topc)).
fof(t52_pre_topc, conjecture, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v4_pre_topc(X2,X1)=>k2_pre_topc(X1,X2)=X2)&((v2_pre_topc(X1)&k2_pre_topc(X1,X2)=X2)=>v4_pre_topc(X2,X1))))), file('data/MPTP/proofs/35d128e821994ca6bf5ed5300c5d736a.E_input', t52_pre_topc)).
fof(c_0_9, plain, ![X5, X6, X7]:(~m1_subset_1(X6,k1_zfmisc_1(X5))|(~r2_hidden(X7,X6)|r2_hidden(X7,X5))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
fof(c_0_10, plain, ![X8, X9]:(~l1_pre_topc(X8)|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X8)))|m1_subset_1(k2_pre_topc(X8,X9),k1_zfmisc_1(u1_struct_0(X8)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_pre_topc])])).
fof(c_0_11, plain, ![X10, X11, X12, X13]:((~r2_hidden(X12,k2_pre_topc(X10,X11))|(~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X10)))|(~v4_pre_topc(X13,X10)|~r1_tarski(X11,X13)|r2_hidden(X12,X13)))|~r2_hidden(X12,u1_struct_0(X10))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~l1_pre_topc(X10))&((m1_subset_1(esk1_3(X10,X11,X12),k1_zfmisc_1(u1_struct_0(X10)))|r2_hidden(X12,k2_pre_topc(X10,X11))|~r2_hidden(X12,u1_struct_0(X10))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~l1_pre_topc(X10))&(((v4_pre_topc(esk1_3(X10,X11,X12),X10)|r2_hidden(X12,k2_pre_topc(X10,X11))|~r2_hidden(X12,u1_struct_0(X10))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~l1_pre_topc(X10))&(r1_tarski(X11,esk1_3(X10,X11,X12))|r2_hidden(X12,k2_pre_topc(X10,X11))|~r2_hidden(X12,u1_struct_0(X10))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~l1_pre_topc(X10)))&(~r2_hidden(X12,esk1_3(X10,X11,X12))|r2_hidden(X12,k2_pre_topc(X10,X11))|~r2_hidden(X12,u1_struct_0(X10))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~l1_pre_topc(X10))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t45_pre_topc])])])])])).
cnf(c_0_12, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (m1_subset_1(k2_pre_topc(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,k2_pre_topc(X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v4_pre_topc(X4,X2)|~r1_tarski(X3,X4)|~r2_hidden(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (r2_hidden(X1,u1_struct_0(X2))|~l1_pre_topc(X2)|~r2_hidden(X1,k2_pre_topc(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(spm,[status(thm)],[c_0_12, c_0_13])).
fof(c_0_16, plain, ![X15, X16, X17, X18, X19]:((~r1_tarski(X15,X16)|(~r2_hidden(X17,X15)|r2_hidden(X17,X16)))&((r2_hidden(esk2_2(X18,X19),X18)|r1_tarski(X18,X19))&(~r2_hidden(esk2_2(X18,X19),X19)|r1_tarski(X18,X19)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_17, plain, ![X21, X22]:((m1_subset_1(esk3_2(X21,X22),k1_zfmisc_1(u1_struct_0(X21)))|v4_pre_topc(k6_setfam_1(u1_struct_0(X21),X22),X21)|~m1_subset_1(X22,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X21))))|(~v2_pre_topc(X21)|~l1_pre_topc(X21)))&((r2_hidden(esk3_2(X21,X22),X22)|v4_pre_topc(k6_setfam_1(u1_struct_0(X21),X22),X21)|~m1_subset_1(X22,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X21))))|(~v2_pre_topc(X21)|~l1_pre_topc(X21)))&(~v4_pre_topc(esk3_2(X21,X22),X21)|v4_pre_topc(k6_setfam_1(u1_struct_0(X21),X22),X21)|~m1_subset_1(X22,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X21))))|(~v2_pre_topc(X21)|~l1_pre_topc(X21))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t44_pre_topc])])])])])).
fof(c_0_18, plain, ![X24, X25, X27]:((m1_subset_1(esk4_2(X24,X25),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X24))))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|(~v2_pre_topc(X24)|~l1_pre_topc(X24)))&((((v4_pre_topc(X27,X24)|~r2_hidden(X27,esk4_2(X24,X25))|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X24)))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|(~v2_pre_topc(X24)|~l1_pre_topc(X24)))&(r1_tarski(X25,X27)|~r2_hidden(X27,esk4_2(X24,X25))|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X24)))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|(~v2_pre_topc(X24)|~l1_pre_topc(X24))))&(~v4_pre_topc(X27,X24)|~r1_tarski(X25,X27)|r2_hidden(X27,esk4_2(X24,X25))|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X24)))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|(~v2_pre_topc(X24)|~l1_pre_topc(X24))))&(k2_pre_topc(X24,X25)=k6_setfam_1(u1_struct_0(X24),esk4_2(X24,X25))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|(~v2_pre_topc(X24)|~l1_pre_topc(X24))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t46_pre_topc])])])])])).
fof(c_0_19, plain, ![X28, X29]:(((r1_tarski(X28,X29)|X28!=X29)&(r1_tarski(X29,X28)|X28!=X29))&(~r1_tarski(X28,X29)|~r1_tarski(X29,X28)|X28=X29)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_20, plain, ![X30, X31]:(~l1_pre_topc(X30)|(~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))|r1_tarski(X31,k2_pre_topc(X30,X31)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t48_pre_topc])])])).
cnf(c_0_21, plain, (r2_hidden(X1,X2)|~r1_tarski(X3,X2)|~v4_pre_topc(X2,X4)|~l1_pre_topc(X4)|~r2_hidden(X1,k2_pre_topc(X4,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X4)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X4)))), inference(csr,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_22, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (r2_hidden(esk3_2(X1,X2),X2)|v4_pre_topc(k6_setfam_1(u1_struct_0(X1),X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_25, negated_conjecture, ~(![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v4_pre_topc(X2,X1)=>k2_pre_topc(X1,X2)=X2)&((v2_pre_topc(X1)&k2_pre_topc(X1,X2)=X2)=>v4_pre_topc(X2,X1)))))), inference(assume_negation,[status(cth)],[t52_pre_topc])).
cnf(c_0_26, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (r1_tarski(X2,k2_pre_topc(X1,X2))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (r1_tarski(X1,X2)|~r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_29, plain, (r1_tarski(k2_pre_topc(X1,X2),X3)|r2_hidden(esk2_2(k2_pre_topc(X1,X2),X3),X4)|~r1_tarski(X2,X4)|~v4_pre_topc(X4,X1)|~l1_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_30, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_31, plain, (v4_pre_topc(X1,X2)|~r2_hidden(X1,esk4_2(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_32, plain, (v4_pre_topc(k6_setfam_1(u1_struct_0(X1),esk4_2(X1,X2)),X1)|r2_hidden(esk3_2(X1,esk4_2(X1,X2)),esk4_2(X1,X2))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_33, negated_conjecture, (l1_pre_topc(esk5_0)&(m1_subset_1(esk6_0,k1_zfmisc_1(u1_struct_0(esk5_0)))&((((v2_pre_topc(esk5_0)|v4_pre_topc(esk6_0,esk5_0))&(k2_pre_topc(esk5_0,esk6_0)=esk6_0|v4_pre_topc(esk6_0,esk5_0)))&(~v4_pre_topc(esk6_0,esk5_0)|v4_pre_topc(esk6_0,esk5_0)))&(((v2_pre_topc(esk5_0)|k2_pre_topc(esk5_0,esk6_0)!=esk6_0)&(k2_pre_topc(esk5_0,esk6_0)=esk6_0|k2_pre_topc(esk5_0,esk6_0)!=esk6_0))&(~v4_pre_topc(esk6_0,esk5_0)|k2_pre_topc(esk5_0,esk6_0)!=esk6_0))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])])).
cnf(c_0_34, plain, (k2_pre_topc(X1,X2)=X2|~r1_tarski(k2_pre_topc(X1,X2),X2)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_35, plain, (r1_tarski(k2_pre_topc(X1,X2),X3)|~r1_tarski(X2,X3)|~v4_pre_topc(X3,X1)|~l1_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_36, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_30])).
cnf(c_0_37, plain, (v4_pre_topc(k6_setfam_1(u1_struct_0(X1),X2),X1)|~v4_pre_topc(esk3_2(X1,X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_38, plain, (v4_pre_topc(k6_setfam_1(u1_struct_0(X1),esk4_2(X1,X2)),X1)|v4_pre_topc(esk3_2(X1,esk4_2(X1,X2)),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(esk3_2(X1,esk4_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_39, plain, (m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|v4_pre_topc(k6_setfam_1(u1_struct_0(X1),X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_40, plain, (k2_pre_topc(X1,X2)=k6_setfam_1(u1_struct_0(X1),esk4_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_41, negated_conjecture, (~v4_pre_topc(esk6_0,esk5_0)|k2_pre_topc(esk5_0,esk6_0)!=esk6_0), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, plain, (k2_pre_topc(X1,X2)=X2|~v4_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])])).
cnf(c_0_43, negated_conjecture, (l1_pre_topc(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_44, negated_conjecture, (m1_subset_1(esk6_0,k1_zfmisc_1(u1_struct_0(esk5_0)))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_45, plain, (v4_pre_topc(k6_setfam_1(u1_struct_0(X1),esk4_2(X1,X2)),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(esk3_2(X1,esk4_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_24])).
cnf(c_0_46, plain, (v4_pre_topc(k2_pre_topc(X1,X2),X1)|m1_subset_1(esk3_2(X1,esk4_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_24])).
cnf(c_0_47, negated_conjecture, (k2_pre_topc(esk5_0,esk6_0)=esk6_0|v4_pre_topc(esk6_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_48, negated_conjecture, (~v4_pre_topc(esk6_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), c_0_44])])).
cnf(c_0_49, negated_conjecture, (v2_pre_topc(esk5_0)|v4_pre_topc(esk6_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_50, plain, (v4_pre_topc(k2_pre_topc(X1,X2),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_40]), c_0_46])).
cnf(c_0_51, negated_conjecture, (k2_pre_topc(esk5_0,esk6_0)=esk6_0), inference(sr,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_52, negated_conjecture, (v2_pre_topc(esk5_0)), inference(sr,[status(thm)],[c_0_49, c_0_48])).
cnf(c_0_53, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52]), c_0_43]), c_0_44])]), c_0_48]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 54
# Proof object clause steps            : 35
# Proof object formula steps           : 19
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 28
# Processed clauses                    : 7784
# ...of these trivial                  : 0
# ...subsumed                          : 6681
# ...remaining for further processing  : 1103
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 86
# Backward-rewritten                   : 4
# Generated clauses                    : 28591
# ...of the previous two non-trivial   : 26806
# Contextual simplify-reflections      : 53
# Paramodulations                      : 28586
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 981
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 974
# Current number of unprocessed clauses: 18948
# ...number of literals in the above   : 225302
# Current number of archived formulas  : 0
# Current number of archived clauses   : 120
# Clause-clause subsumption calls (NU) : 358170
# Rec. Clause-clause subsumption calls : 24133
# Non-unit clause-clause subsumptions  : 3914
# Unit Clause-clause subsumption calls : 24
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1349791

# -------------------------------------------------
# User time                : 0.879 s
# System time              : 0.010 s
# Total time               : 0.889 s
# Maximum resident set size: 28760 pages
