# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B02_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_wellord2, axiom, ![X1, X2]:(v1_relat_1(X2)=>(X2=k1_wellord2(X1)<=>(k3_relat_1(X2)=X1&![X3, X4]:((r2_hidden(X3,X1)&r2_hidden(X4,X1))=>(r2_hidden(k4_tarski(X3,X4),X2)<=>r1_tarski(X3,X4)))))), file('data/MPTP/proofs/361ac4f200534ff38d748e516b065f49.E_input', d1_wellord2)).
fof(dt_k1_wellord2, axiom, ![X1]:v1_relat_1(k1_wellord2(X1)), file('data/MPTP/proofs/361ac4f200534ff38d748e516b065f49.E_input', dt_k1_wellord2)).
fof(l4_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v6_relat_2(X1)<=>![X2, X3]:~((r2_hidden(X2,k3_relat_1(X1))&(r2_hidden(X3,k3_relat_1(X1))&(~(X2=X3)&(~(r2_hidden(k4_tarski(X2,X3),X1))&~(r2_hidden(k4_tarski(X3,X2),X1))))))))), file('data/MPTP/proofs/361ac4f200534ff38d748e516b065f49.E_input', l4_wellord1)).
fof(redefinition_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)<=>r1_tarski(X1,X2))), file('data/MPTP/proofs/361ac4f200534ff38d748e516b065f49.E_input', redefinition_r1_ordinal1)).
fof(t26_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r1_ordinal1(X1,X2)|r2_hidden(X2,X1)))), file('data/MPTP/proofs/361ac4f200534ff38d748e516b065f49.E_input', t26_ordinal1)).
fof(d2_ordinal1, axiom, ![X1]:(v1_ordinal1(X1)<=>![X2]:(r2_hidden(X2,X1)=>r1_tarski(X2,X1))), file('data/MPTP/proofs/361ac4f200534ff38d748e516b065f49.E_input', d2_ordinal1)).
fof(t23_ordinal1, axiom, ![X1, X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)=>v3_ordinal1(X1))), file('data/MPTP/proofs/361ac4f200534ff38d748e516b065f49.E_input', t23_ordinal1)).
fof(cc1_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>(v1_ordinal1(X1)&v2_ordinal1(X1))), file('data/MPTP/proofs/361ac4f200534ff38d748e516b065f49.E_input', cc1_ordinal1)).
fof(t4_wellord2, conjecture, ![X1]:(v3_ordinal1(X1)=>v6_relat_2(k1_wellord2(X1))), file('data/MPTP/proofs/361ac4f200534ff38d748e516b065f49.E_input', t4_wellord2)).
fof(c_0_9, plain, ![X6, X7, X8, X9]:(((k3_relat_1(X7)=X6|X7!=k1_wellord2(X6)|~v1_relat_1(X7))&((~r2_hidden(k4_tarski(X8,X9),X7)|r1_tarski(X8,X9)|(~r2_hidden(X8,X6)|~r2_hidden(X9,X6))|X7!=k1_wellord2(X6)|~v1_relat_1(X7))&(~r1_tarski(X8,X9)|r2_hidden(k4_tarski(X8,X9),X7)|(~r2_hidden(X8,X6)|~r2_hidden(X9,X6))|X7!=k1_wellord2(X6)|~v1_relat_1(X7))))&(((r2_hidden(esk2_2(X6,X7),X6)|k3_relat_1(X7)!=X6|X7=k1_wellord2(X6)|~v1_relat_1(X7))&(r2_hidden(esk3_2(X6,X7),X6)|k3_relat_1(X7)!=X6|X7=k1_wellord2(X6)|~v1_relat_1(X7)))&((~r2_hidden(k4_tarski(esk2_2(X6,X7),esk3_2(X6,X7)),X7)|~r1_tarski(esk2_2(X6,X7),esk3_2(X6,X7))|k3_relat_1(X7)!=X6|X7=k1_wellord2(X6)|~v1_relat_1(X7))&(r2_hidden(k4_tarski(esk2_2(X6,X7),esk3_2(X6,X7)),X7)|r1_tarski(esk2_2(X6,X7),esk3_2(X6,X7))|k3_relat_1(X7)!=X6|X7=k1_wellord2(X6)|~v1_relat_1(X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_wellord2])])])])])).
fof(c_0_10, plain, ![X19]:v1_relat_1(k1_wellord2(X19)), inference(variable_rename,[status(thm)],[dt_k1_wellord2])).
fof(c_0_11, plain, ![X24, X25, X26]:((~v6_relat_2(X24)|(~r2_hidden(X25,k3_relat_1(X24))|(~r2_hidden(X26,k3_relat_1(X24))|(X25=X26|(r2_hidden(k4_tarski(X25,X26),X24)|r2_hidden(k4_tarski(X26,X25),X24)))))|~v1_relat_1(X24))&((r2_hidden(esk5_1(X24),k3_relat_1(X24))|v6_relat_2(X24)|~v1_relat_1(X24))&((r2_hidden(esk6_1(X24),k3_relat_1(X24))|v6_relat_2(X24)|~v1_relat_1(X24))&((esk5_1(X24)!=esk6_1(X24)|v6_relat_2(X24)|~v1_relat_1(X24))&((~r2_hidden(k4_tarski(esk5_1(X24),esk6_1(X24)),X24)|v6_relat_2(X24)|~v1_relat_1(X24))&(~r2_hidden(k4_tarski(esk6_1(X24),esk5_1(X24)),X24)|v6_relat_2(X24)|~v1_relat_1(X24))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l4_wellord1])])])])])])])).
cnf(c_0_12, plain, (k3_relat_1(X1)=X2|X1!=k1_wellord2(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (v1_relat_1(k1_wellord2(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (r2_hidden(esk5_1(X1),k3_relat_1(X1))|v6_relat_2(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (k3_relat_1(k1_wellord2(X1))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_12]), c_0_13])])).
cnf(c_0_16, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r1_tarski(X1,X2)|~r2_hidden(X1,X4)|~r2_hidden(X2,X4)|X3!=k1_wellord2(X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, plain, (r2_hidden(esk6_1(X1),k3_relat_1(X1))|v6_relat_2(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, plain, (r2_hidden(esk5_1(k1_wellord2(X1)),X1)|v6_relat_2(k1_wellord2(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_13])])).
fof(c_0_19, plain, ![X12, X13]:((~r1_ordinal1(X12,X13)|r1_tarski(X12,X13)|(~v3_ordinal1(X12)|~v3_ordinal1(X13)))&(~r1_tarski(X12,X13)|r1_ordinal1(X12,X13)|(~v3_ordinal1(X12)|~v3_ordinal1(X13)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r1_ordinal1])])])).
fof(c_0_20, plain, ![X16, X17]:(~v3_ordinal1(X16)|(~v3_ordinal1(X17)|(r1_ordinal1(X16,X17)|r2_hidden(X17,X16)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t26_ordinal1])])])).
cnf(c_0_21, plain, (v6_relat_2(X1)|~r2_hidden(k4_tarski(esk5_1(X1),esk6_1(X1)),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_22, plain, (r2_hidden(k4_tarski(X1,esk6_1(X2)),X3)|v6_relat_2(X2)|X3!=k1_wellord2(k3_relat_1(X2))|~r1_tarski(X1,esk6_1(X2))|~r2_hidden(X1,k3_relat_1(X2))|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
fof(c_0_23, plain, ![X20, X21, X22]:((~v1_ordinal1(X20)|(~r2_hidden(X21,X20)|r1_tarski(X21,X20)))&((r2_hidden(esk4_1(X22),X22)|v1_ordinal1(X22))&(~r1_tarski(esk4_1(X22),X22)|v1_ordinal1(X22)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_ordinal1])])])])])])).
cnf(c_0_24, plain, (v6_relat_2(X1)|~r2_hidden(k4_tarski(esk6_1(X1),esk5_1(X1)),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_25, plain, (r2_hidden(k4_tarski(X1,esk5_1(k1_wellord2(X2))),X3)|v6_relat_2(k1_wellord2(X2))|X3!=k1_wellord2(X2)|~r1_tarski(X1,esk5_1(k1_wellord2(X2)))|~r2_hidden(X1,X2)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_16, c_0_18])).
cnf(c_0_26, plain, (r2_hidden(esk6_1(k1_wellord2(X1)),X1)|v6_relat_2(k1_wellord2(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_15]), c_0_13])])).
cnf(c_0_27, plain, (r1_tarski(X1,X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, plain, (r1_ordinal1(X1,X2)|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_29, plain, ![X14, X15]:(~v3_ordinal1(X15)|(~r2_hidden(X14,X15)|v3_ordinal1(X14))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_ordinal1])])).
cnf(c_0_30, plain, (v6_relat_2(X1)|k1_wellord2(k3_relat_1(X1))!=X1|~r1_tarski(esk5_1(X1),esk6_1(X1))|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_14])).
cnf(c_0_31, plain, (r1_tarski(X2,X1)|~v1_ordinal1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, plain, (v6_relat_2(k1_wellord2(X1))|~r1_tarski(esk6_1(k1_wellord2(X1)),esk5_1(k1_wellord2(X1)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_13])]), c_0_26])).
cnf(c_0_33, plain, (r1_tarski(X1,X2)|r2_hidden(X2,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
fof(c_0_34, plain, ![X18]:((v1_ordinal1(X18)|~v3_ordinal1(X18))&(v2_ordinal1(X18)|~v3_ordinal1(X18))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_ordinal1])])])).
cnf(c_0_35, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_36, negated_conjecture, ~(![X1]:(v3_ordinal1(X1)=>v6_relat_2(k1_wellord2(X1)))), inference(assume_negation,[status(cth)],[t4_wellord2])).
cnf(c_0_37, plain, (v6_relat_2(X1)|k1_wellord2(k3_relat_1(X1))!=X1|~v1_ordinal1(esk6_1(X1))|~r2_hidden(esk5_1(X1),esk6_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_38, plain, (r2_hidden(esk5_1(k1_wellord2(X1)),esk6_1(k1_wellord2(X1)))|v6_relat_2(k1_wellord2(X1))|~v3_ordinal1(esk5_1(k1_wellord2(X1)))|~v3_ordinal1(esk6_1(k1_wellord2(X1)))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_39, plain, (v1_ordinal1(X1)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_40, plain, (v6_relat_2(X1)|v3_ordinal1(esk5_1(X1))|~v1_relat_1(X1)|~v3_ordinal1(k3_relat_1(X1))), inference(spm,[status(thm)],[c_0_35, c_0_14])).
fof(c_0_41, negated_conjecture, (v3_ordinal1(esk1_0)&~v6_relat_2(k1_wellord2(esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_36])])])).
cnf(c_0_42, plain, (v6_relat_2(k1_wellord2(X1))|~v3_ordinal1(esk5_1(k1_wellord2(X1)))|~v3_ordinal1(esk6_1(k1_wellord2(X1)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_15]), c_0_13])]), c_0_39])).
cnf(c_0_43, plain, (v6_relat_2(X1)|v3_ordinal1(esk6_1(X1))|~v1_relat_1(X1)|~v3_ordinal1(k3_relat_1(X1))), inference(spm,[status(thm)],[c_0_35, c_0_17])).
cnf(c_0_44, plain, (v6_relat_2(k1_wellord2(X1))|v3_ordinal1(esk5_1(k1_wellord2(X1)))|~v3_ordinal1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_15]), c_0_13])])).
cnf(c_0_45, negated_conjecture, (~v6_relat_2(k1_wellord2(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_46, plain, (v6_relat_2(k1_wellord2(X1))|~v3_ordinal1(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_13]), c_0_15])]), c_0_44])).
cnf(c_0_47, negated_conjecture, (v3_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_48, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 49
# Proof object clause steps            : 30
# Proof object formula steps           : 19
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 25
# Processed clauses                    : 76
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 73
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 0
# Generated clauses                    : 61
# ...of the previous two non-trivial   : 50
# Contextual simplify-reflections      : 4
# Paramodulations                      : 59
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 45
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 41
# Current number of unprocessed clauses: 24
# ...number of literals in the above   : 197
# Current number of archived formulas  : 0
# Current number of archived clauses   : 28
# Clause-clause subsumption calls (NU) : 897
# Rec. Clause-clause subsumption calls : 283
# Non-unit clause-clause subsumptions  : 10
# Unit Clause-clause subsumption calls : 7
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3515

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.004 s
# Total time               : 0.011 s
# Maximum resident set size: 28680 pages
