# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t36_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>r1_lattice3(X1,k4_waybel_0(X1,X2),X3))))), file('data/MPTP/proofs/366bd076f9b0455e8c0ac6dfef428f0a.E_input', t36_waybel_0)).
fof(t15_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k4_waybel_0(X1,X2)=a_2_1_waybel_0(X1,X2))), file('data/MPTP/proofs/366bd076f9b0455e8c0ac6dfef428f0a.E_input', t15_waybel_0)).
fof(t16_waybel_0, axiom, ![X1]:((v3_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(r1_tarski(X2,k3_waybel_0(X1,X2))&r1_tarski(X2,k4_waybel_0(X1,X2))))), file('data/MPTP/proofs/366bd076f9b0455e8c0ac6dfef428f0a.E_input', t16_waybel_0)).
fof(t9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(r1_tarski(X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X3,X4)=>r1_lattice3(X1,X2,X4))&(r2_lattice3(X1,X3,X4)=>r2_lattice3(X1,X2,X4)))))), file('data/MPTP/proofs/366bd076f9b0455e8c0ac6dfef428f0a.E_input', t9_yellow_0)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/366bd076f9b0455e8c0ac6dfef428f0a.E_input', d8_lattice3)).
fof(t26_orders_2, axiom, ![X1]:((v4_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X3,X4))=>r1_orders_2(X1,X2,X4)))))), file('data/MPTP/proofs/366bd076f9b0455e8c0ac6dfef428f0a.E_input', t26_orders_2)).
fof(fraenkel_a_2_1_waybel_0, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&l1_orders_2(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r2_hidden(X1,a_2_1_waybel_0(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X5,X4)&r2_hidden(X5,X3))))))), file('data/MPTP/proofs/366bd076f9b0455e8c0ac6dfef428f0a.E_input', fraenkel_a_2_1_waybel_0)).
fof(c_0_7, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>r1_lattice3(X1,k4_waybel_0(X1,X2),X3)))))), inference(assume_negation,[status(cth)],[t36_waybel_0])).
fof(c_0_8, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&((~r1_lattice3(esk1_0,esk2_0,esk3_0)|~r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0))&(r1_lattice3(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_7])])])])).
fof(c_0_9, plain, ![X29, X30]:(v2_struct_0(X29)|~l1_orders_2(X29)|(~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))|k4_waybel_0(X29,X30)=a_2_1_waybel_0(X29,X30))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t15_waybel_0])])])])).
fof(c_0_10, plain, ![X31, X32]:((r1_tarski(X32,k3_waybel_0(X31,X32))|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|(~v3_orders_2(X31)|~l1_orders_2(X31)))&(r1_tarski(X32,k4_waybel_0(X31,X32))|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|(~v3_orders_2(X31)|~l1_orders_2(X31)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_waybel_0])])])])).
fof(c_0_11, plain, ![X13, X14, X15, X16]:((~r1_lattice3(X13,X15,X16)|r1_lattice3(X13,X14,X16)|~m1_subset_1(X16,u1_struct_0(X13))|~r1_tarski(X14,X15)|~l1_orders_2(X13))&(~r2_lattice3(X13,X15,X16)|r2_lattice3(X13,X14,X16)|~m1_subset_1(X16,u1_struct_0(X13))|~r1_tarski(X14,X15)|~l1_orders_2(X13))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_yellow_0])])])])).
cnf(c_0_12, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (v2_struct_0(X1)|k4_waybel_0(X1,X2)=a_2_1_waybel_0(X1,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_17, plain, (r1_tarski(X1,k4_waybel_0(X2,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v3_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_19, plain, (r1_lattice3(X1,X4,X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_tarski(X4,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20, negated_conjecture, (r1_lattice3(esk1_0,a_2_1_waybel_0(esk1_0,esk2_0),esk3_0)|r1_lattice3(esk1_0,esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_21, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_22, negated_conjecture, (r1_tarski(X1,k4_waybel_0(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_15])])).
fof(c_0_23, plain, ![X24, X25, X26, X27]:((~r1_lattice3(X24,X25,X26)|(~m1_subset_1(X27,u1_struct_0(X24))|(~r2_hidden(X27,X25)|r1_orders_2(X24,X26,X27)))|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))&((m1_subset_1(esk6_3(X24,X25,X26),u1_struct_0(X24))|r1_lattice3(X24,X25,X26)|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))&((r2_hidden(esk6_3(X24,X25,X26),X25)|r1_lattice3(X24,X25,X26)|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))&(~r1_orders_2(X24,X26,esk6_3(X24,X25,X26))|r1_lattice3(X24,X25,X26)|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
cnf(c_0_24, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,X1,esk3_0)|~r1_tarski(X1,a_2_1_waybel_0(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_15])])).
cnf(c_0_25, negated_conjecture, (r1_tarski(X1,a_2_1_waybel_0(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_13]), c_0_15])]), c_0_16])).
fof(c_0_26, plain, ![X9, X10, X11, X12]:(~v4_orders_2(X9)|~l1_orders_2(X9)|(~m1_subset_1(X10,u1_struct_0(X9))|(~m1_subset_1(X11,u1_struct_0(X9))|(~m1_subset_1(X12,u1_struct_0(X9))|(~r1_orders_2(X9,X10,X11)|~r1_orders_2(X9,X11,X12)|r1_orders_2(X9,X10,X12)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t26_orders_2])])])).
fof(c_0_27, plain, ![X17, X18, X19, X22, X23]:(((m1_subset_1(esk4_3(X17,X18,X19),u1_struct_0(X18))|~r2_hidden(X17,a_2_1_waybel_0(X18,X19))|(v2_struct_0(X18)|~l1_orders_2(X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))&((X17=esk4_3(X17,X18,X19)|~r2_hidden(X17,a_2_1_waybel_0(X18,X19))|(v2_struct_0(X18)|~l1_orders_2(X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))&((m1_subset_1(esk5_3(X17,X18,X19),u1_struct_0(X18))|~r2_hidden(X17,a_2_1_waybel_0(X18,X19))|(v2_struct_0(X18)|~l1_orders_2(X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))&((r1_orders_2(X18,esk5_3(X17,X18,X19),esk4_3(X17,X18,X19))|~r2_hidden(X17,a_2_1_waybel_0(X18,X19))|(v2_struct_0(X18)|~l1_orders_2(X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))&(r2_hidden(esk5_3(X17,X18,X19),X19)|~r2_hidden(X17,a_2_1_waybel_0(X18,X19))|(v2_struct_0(X18)|~l1_orders_2(X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))))))&(~m1_subset_1(X22,u1_struct_0(X18))|(X17!=X22|(~m1_subset_1(X23,u1_struct_0(X18))|(~r1_orders_2(X18,X23,X22)|~r2_hidden(X23,X19))))|r2_hidden(X17,a_2_1_waybel_0(X18,X19))|(v2_struct_0(X18)|~l1_orders_2(X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_1_waybel_0])])])])])])).
cnf(c_0_28, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_14])])).
cnf(c_0_30, plain, (r1_orders_2(X1,X2,X4)|~v4_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_31, plain, (r1_orders_2(X1,esk5_3(X2,X1,X3),esk4_3(X2,X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,a_2_1_waybel_0(X1,X3))|~l1_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_32, plain, (m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_34, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,X1)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_21]), c_0_15])])).
cnf(c_0_35, plain, (r2_hidden(esk5_3(X1,X2,X3),X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36, plain, (r1_orders_2(X1,X2,esk4_3(X3,X1,X4))|v2_struct_0(X1)|~r2_hidden(X3,a_2_1_waybel_0(X1,X4))|~r1_orders_2(X1,X2,esk5_3(X3,X1,X4))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v4_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_37, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk5_3(X1,X2,esk2_0))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,esk2_0))|~m1_subset_1(esk5_3(X1,X2,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_38, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_39, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk4_3(X1,esk1_0,esk2_0))|~r2_hidden(X1,a_2_1_waybel_0(esk1_0,esk2_0))|~m1_subset_1(esk5_3(X1,esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_14]), c_0_21]), c_0_15]), c_0_38])]), c_0_16])).
cnf(c_0_40, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk4_3(X1,esk1_0,esk2_0))|~r2_hidden(X1,a_2_1_waybel_0(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_32]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_41, plain, (X1=esk4_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_42, negated_conjecture, (~r1_lattice3(esk1_0,esk2_0,esk3_0)|~r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_43, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,X1)|~r2_hidden(X1,a_2_1_waybel_0(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_44, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_45, negated_conjecture, (~r1_lattice3(esk1_0,a_2_1_waybel_0(esk1_0,esk2_0),esk3_0)|~r1_lattice3(esk1_0,esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_13]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_46, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk6_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_47, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk6_3(X1,a_2_1_waybel_0(esk1_0,esk2_0),X2))|r1_lattice3(X1,a_2_1_waybel_0(esk1_0,esk2_0),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_48, negated_conjecture, (~r1_lattice3(esk1_0,a_2_1_waybel_0(esk1_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_29])])).
cnf(c_0_49, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_21]), c_0_15])]), c_0_48]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 50
# Proof object clause steps            : 35
# Proof object formula steps           : 15
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 43
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 24
# Processed clauses                    : 102
# ...of these trivial                  : 0
# ...subsumed                          : 6
# ...remaining for further processing  : 96
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 5
# Generated clauses                    : 88
# ...of the previous two non-trivial   : 83
# Contextual simplify-reflections      : 2
# Paramodulations                      : 87
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 60
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 51
# Current number of unprocessed clauses: 27
# ...number of literals in the above   : 238
# Current number of archived formulas  : 0
# Current number of archived clauses   : 35
# Clause-clause subsumption calls (NU) : 1834
# Rec. Clause-clause subsumption calls : 94
# Non-unit clause-clause subsumptions  : 13
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6651

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.003 s
# Total time               : 0.014 s
# Maximum resident set size: 28624 pages
