# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t24_yellow_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1)), file('data/MPTP/proofs/37231a32a4e348b69f500c44a0b8d1b8.E_input', t24_yellow_1)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/37231a32a4e348b69f500c44a0b8d1b8.E_input', t1_xboole_1)).
fof(l49_zfmisc_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>r1_tarski(X1,k3_tarski(X2))), file('data/MPTP/proofs/37231a32a4e348b69f500c44a0b8d1b8.E_input', l49_zfmisc_1)).
fof(d1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1))))))))), file('data/MPTP/proofs/37231a32a4e348b69f500c44a0b8d1b8.E_input', d1_pre_topc)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/37231a32a4e348b69f500c44a0b8d1b8.E_input', d10_xboole_0)).
fof(t95_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>r1_tarski(k3_tarski(X1),k3_tarski(X2))), file('data/MPTP/proofs/37231a32a4e348b69f500c44a0b8d1b8.E_input', t95_zfmisc_1)).
fof(t99_zfmisc_1, axiom, ![X1]:k3_tarski(k1_zfmisc_1(X1))=X1, file('data/MPTP/proofs/37231a32a4e348b69f500c44a0b8d1b8.E_input', t99_zfmisc_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/37231a32a4e348b69f500c44a0b8d1b8.E_input', t2_subset)).
fof(dt_u1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/37231a32a4e348b69f500c44a0b8d1b8.E_input', dt_u1_pre_topc)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/37231a32a4e348b69f500c44a0b8d1b8.E_input', fc1_subset_1)).
fof(t14_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>(r2_hidden(k3_tarski(X1),X1)=>k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1))), file('data/MPTP/proofs/37231a32a4e348b69f500c44a0b8d1b8.E_input', t14_yellow_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/37231a32a4e348b69f500c44a0b8d1b8.E_input', d1_xboole_0)).
fof(c_0_12, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1))), inference(assume_negation,[status(cth)],[t24_yellow_1])).
fof(c_0_13, plain, ![X5, X6, X7]:(~r1_tarski(X5,X6)|~r1_tarski(X6,X7)|r1_tarski(X5,X7)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_14, plain, ![X8, X9]:(~r2_hidden(X8,X9)|r1_tarski(X8,k3_tarski(X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l49_zfmisc_1])])).
fof(c_0_15, plain, ![X10, X11, X12, X13]:(((r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|~v2_pre_topc(X10)|~l1_pre_topc(X10))&((~m1_subset_1(X11,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|(~r1_tarski(X11,u1_pre_topc(X10))|r2_hidden(k5_setfam_1(u1_struct_0(X10),X11),u1_pre_topc(X10)))|~v2_pre_topc(X10)|~l1_pre_topc(X10))&(~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X10)))|(~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X10)))|(~r2_hidden(X12,u1_pre_topc(X10))|~r2_hidden(X13,u1_pre_topc(X10))|r2_hidden(k9_subset_1(u1_struct_0(X10),X12,X13),u1_pre_topc(X10))))|~v2_pre_topc(X10)|~l1_pre_topc(X10))))&(((m1_subset_1(esk3_1(X10),k1_zfmisc_1(u1_struct_0(X10)))|m1_subset_1(esk2_1(X10),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&((m1_subset_1(esk4_1(X10),k1_zfmisc_1(u1_struct_0(X10)))|m1_subset_1(esk2_1(X10),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&(((r2_hidden(esk3_1(X10),u1_pre_topc(X10))|m1_subset_1(esk2_1(X10),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&(r2_hidden(esk4_1(X10),u1_pre_topc(X10))|m1_subset_1(esk2_1(X10),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10)))&(~r2_hidden(k9_subset_1(u1_struct_0(X10),esk3_1(X10),esk4_1(X10)),u1_pre_topc(X10))|m1_subset_1(esk2_1(X10),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10)))))&(((m1_subset_1(esk3_1(X10),k1_zfmisc_1(u1_struct_0(X10)))|r1_tarski(esk2_1(X10),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&((m1_subset_1(esk4_1(X10),k1_zfmisc_1(u1_struct_0(X10)))|r1_tarski(esk2_1(X10),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&(((r2_hidden(esk3_1(X10),u1_pre_topc(X10))|r1_tarski(esk2_1(X10),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&(r2_hidden(esk4_1(X10),u1_pre_topc(X10))|r1_tarski(esk2_1(X10),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10)))&(~r2_hidden(k9_subset_1(u1_struct_0(X10),esk3_1(X10),esk4_1(X10)),u1_pre_topc(X10))|r1_tarski(esk2_1(X10),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10)))))&((m1_subset_1(esk3_1(X10),k1_zfmisc_1(u1_struct_0(X10)))|~r2_hidden(k5_setfam_1(u1_struct_0(X10),esk2_1(X10)),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&((m1_subset_1(esk4_1(X10),k1_zfmisc_1(u1_struct_0(X10)))|~r2_hidden(k5_setfam_1(u1_struct_0(X10),esk2_1(X10)),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&(((r2_hidden(esk3_1(X10),u1_pre_topc(X10))|~r2_hidden(k5_setfam_1(u1_struct_0(X10),esk2_1(X10)),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&(r2_hidden(esk4_1(X10),u1_pre_topc(X10))|~r2_hidden(k5_setfam_1(u1_struct_0(X10),esk2_1(X10)),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10)))&(~r2_hidden(k9_subset_1(u1_struct_0(X10),esk3_1(X10),esk4_1(X10)),u1_pre_topc(X10))|~r2_hidden(k5_setfam_1(u1_struct_0(X10),esk2_1(X10)),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])).
fof(c_0_16, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&k4_yellow_0(k2_yellow_1(u1_pre_topc(esk1_0)))!=u1_struct_0(esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])).
cnf(c_0_17, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, plain, (r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, plain, ![X17, X18]:(((r1_tarski(X17,X18)|X17!=X18)&(r1_tarski(X18,X17)|X17!=X18))&(~r1_tarski(X17,X18)|~r1_tarski(X18,X17)|X17=X18)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_23, plain, (r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X3,X2)|~r1_tarski(X1,X3)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_24, negated_conjecture, (r2_hidden(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])])).
fof(c_0_25, plain, ![X19, X20]:(~r1_tarski(X19,X20)|r1_tarski(k3_tarski(X19),k3_tarski(X20))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t95_zfmisc_1])])).
fof(c_0_26, plain, ![X21]:k3_tarski(k1_zfmisc_1(X21))=X21, inference(variable_rename,[status(thm)],[t99_zfmisc_1])).
fof(c_0_27, plain, ![X22, X23]:(~m1_subset_1(X22,X23)|(v1_xboole_0(X23)|r2_hidden(X22,X23))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_28, plain, ![X24]:(~l1_pre_topc(X24)|m1_subset_1(u1_pre_topc(X24),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X24))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_pre_topc])])).
fof(c_0_29, plain, ![X25]:~v1_xboole_0(k1_zfmisc_1(X25)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
cnf(c_0_30, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, negated_conjecture, (r1_tarski(X1,k3_tarski(u1_pre_topc(esk1_0)))|~r1_tarski(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_32, plain, (r1_tarski(k3_tarski(X1),k3_tarski(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (k3_tarski(k1_zfmisc_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, negated_conjecture, (k3_tarski(u1_pre_topc(esk1_0))=X1|~r1_tarski(k3_tarski(u1_pre_topc(esk1_0)),X1)|~r1_tarski(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_38, plain, (r1_tarski(k3_tarski(X1),X2)|~r1_tarski(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_39, plain, (r2_hidden(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])).
fof(c_0_40, plain, ![X26]:(v1_xboole_0(X26)|(~r2_hidden(k3_tarski(X26),X26)|k4_yellow_0(k2_yellow_1(X26))=k3_tarski(X26))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_yellow_1])])])).
fof(c_0_41, plain, ![X27, X28, X29]:((~v1_xboole_0(X27)|~r2_hidden(X28,X27))&(r2_hidden(esk5_1(X29),X29)|v1_xboole_0(X29))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_42, negated_conjecture, (k3_tarski(u1_pre_topc(esk1_0))=X1|~r1_tarski(u1_pre_topc(esk1_0),k1_zfmisc_1(X1))|~r1_tarski(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_43, plain, (r1_tarski(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r1_tarski(X1,u1_pre_topc(X2))|~l1_pre_topc(X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_39]), c_0_33])).
cnf(c_0_44, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_45, plain, (v1_xboole_0(X1)|k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_46, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_47, negated_conjecture, (k3_tarski(u1_pre_topc(esk1_0))=u1_struct_0(X1)|~r1_tarski(u1_struct_0(X1),u1_struct_0(esk1_0))|~r1_tarski(u1_pre_topc(esk1_0),u1_pre_topc(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_48, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_49, plain, (k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(csr,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_50, negated_conjecture, (k3_tarski(u1_pre_topc(esk1_0))=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_48]), c_0_21])])).
cnf(c_0_51, negated_conjecture, (k4_yellow_0(k2_yellow_1(u1_pre_topc(esk1_0)))!=u1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_52, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_24])]), c_0_51]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 53
# Proof object clause steps            : 28
# Proof object formula steps           : 25
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 35
# Processed clauses                    : 84
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 82
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 4
# Generated clauses                    : 71
# ...of the previous two non-trivial   : 66
# Contextual simplify-reflections      : 1
# Paramodulations                      : 69
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 41
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 31
# Current number of unprocessed clauses: 44
# ...number of literals in the above   : 181
# Current number of archived formulas  : 0
# Current number of archived clauses   : 39
# Clause-clause subsumption calls (NU) : 414
# Rec. Clause-clause subsumption calls : 82
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4059

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28732 pages
