# No SInE strategy applied
# Auto-Mode selected heuristic H_____011_C18_F1_PI_SE_SP_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4_yellow_1, axiom, ![X1]:k3_yellow_1(X1)=k2_yellow_1(k9_setfam_1(X1)), file('data/MPTP/proofs/373ff35f98834661965d1765faad6bb8.E_input', t4_yellow_1)).
fof(redefinition_k9_setfam_1, axiom, ![X1]:k9_setfam_1(X1)=k1_zfmisc_1(X1), file('data/MPTP/proofs/373ff35f98834661965d1765faad6bb8.E_input', redefinition_k9_setfam_1)).
fof(t2_yellow19, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(X1)))&(v2_waybel_0(X2,k3_yellow_1(X1))&(v13_waybel_0(X2,k3_yellow_1(X1))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X1))))))))=>![X3]:~((r2_hidden(X3,X2)&v1_xboole_0(X3))))), file('data/MPTP/proofs/373ff35f98834661965d1765faad6bb8.E_input', t2_yellow19)).
fof(t11_waybel_7, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X1))))=>(v13_waybel_0(X2,k3_yellow_1(X1))<=>![X3, X4]:((r1_tarski(X3,X4)&(r1_tarski(X4,X1)&r2_hidden(X3,X2)))=>r2_hidden(X4,X2)))), file('data/MPTP/proofs/373ff35f98834661965d1765faad6bb8.E_input', t11_waybel_7)).
fof(t1_yellow_1, axiom, ![X1]:(u1_struct_0(k2_yellow_1(X1))=X1&u1_orders_2(k2_yellow_1(X1))=k1_yellow_1(X1)), file('data/MPTP/proofs/373ff35f98834661965d1765faad6bb8.E_input', t1_yellow_1)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/373ff35f98834661965d1765faad6bb8.E_input', d1_zfmisc_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/373ff35f98834661965d1765faad6bb8.E_input', d3_tarski)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/373ff35f98834661965d1765faad6bb8.E_input', d1_xboole_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/373ff35f98834661965d1765faad6bb8.E_input', t3_subset)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/373ff35f98834661965d1765faad6bb8.E_input', d10_xboole_0)).
fof(d7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(v1_subset_1(X2,X1)<=>~(X2=X1))), file('data/MPTP/proofs/373ff35f98834661965d1765faad6bb8.E_input', d7_subset_1)).
fof(c_0_11, plain, ![X21]:k3_yellow_1(X21)=k2_yellow_1(k9_setfam_1(X21)), inference(variable_rename,[status(thm)],[t4_yellow_1])).
fof(c_0_12, plain, ![X20]:k9_setfam_1(X20)=k1_zfmisc_1(X20), inference(variable_rename,[status(thm)],[redefinition_k9_setfam_1])).
fof(c_0_13, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(X1)))&(v2_waybel_0(X2,k3_yellow_1(X1))&(v13_waybel_0(X2,k3_yellow_1(X1))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X1))))))))=>![X3]:~((r2_hidden(X3,X2)&v1_xboole_0(X3)))))), inference(assume_negation,[status(cth)],[t2_yellow19])).
fof(c_0_14, plain, ![X32, X33, X34, X35]:((~v13_waybel_0(X33,k3_yellow_1(X32))|(~r1_tarski(X34,X35)|(~r1_tarski(X35,X32)|~r2_hidden(X34,X33))|r2_hidden(X35,X33))|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X32)))))&(((r1_tarski(esk7_2(X32,X33),esk8_2(X32,X33))|v13_waybel_0(X33,k3_yellow_1(X32))|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X32)))))&((r1_tarski(esk8_2(X32,X33),X32)|v13_waybel_0(X33,k3_yellow_1(X32))|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X32)))))&(r2_hidden(esk7_2(X32,X33),X33)|v13_waybel_0(X33,k3_yellow_1(X32))|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X32)))))))&(~r2_hidden(esk8_2(X32,X33),X33)|v13_waybel_0(X33,k3_yellow_1(X32))|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X32))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_waybel_7])])])])])).
cnf(c_0_15, plain, (k3_yellow_1(X1)=k2_yellow_1(k9_setfam_1(X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (k9_setfam_1(X1)=k1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_17, negated_conjecture, (~v1_xboole_0(esk1_0)&((~v1_xboole_0(esk2_0)&(v1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(esk1_0)))&(v2_waybel_0(esk2_0,k3_yellow_1(esk1_0))&(v13_waybel_0(esk2_0,k3_yellow_1(esk1_0))&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(esk1_0))))))))&(r2_hidden(esk3_0,esk2_0)&v1_xboole_0(esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_13])])])])).
cnf(c_0_18, plain, (r2_hidden(X4,X1)|~v13_waybel_0(X1,k3_yellow_1(X2))|~r1_tarski(X3,X4)|~r1_tarski(X4,X2)|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (k3_yellow_1(X1)=k2_yellow_1(k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_15, c_0_16])).
fof(c_0_20, plain, ![X10]:(u1_struct_0(k2_yellow_1(X10))=X10&u1_orders_2(k2_yellow_1(X10))=k1_yellow_1(X10)), inference(variable_rename,[status(thm)],[t1_yellow_1])).
cnf(c_0_21, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_22, plain, ![X13, X14, X15, X16, X17, X18]:(((~r2_hidden(X15,X14)|r1_tarski(X15,X13)|X14!=k1_zfmisc_1(X13))&(~r1_tarski(X16,X13)|r2_hidden(X16,X14)|X14!=k1_zfmisc_1(X13)))&((~r2_hidden(esk4_2(X17,X18),X18)|~r1_tarski(esk4_2(X17,X18),X17)|X18=k1_zfmisc_1(X17))&(r2_hidden(esk4_2(X17,X18),X18)|r1_tarski(esk4_2(X17,X18),X17)|X18=k1_zfmisc_1(X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
cnf(c_0_23, plain, (r2_hidden(X4,X1)|~r2_hidden(X3,X1)|~r1_tarski(X4,X2)|~r1_tarski(X3,X4)|~v13_waybel_0(X1,k2_yellow_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k1_zfmisc_1(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_19]), c_0_19])).
cnf(c_0_24, plain, (u1_struct_0(k2_yellow_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k1_zfmisc_1(esk1_0)))))), inference(rw,[status(thm)],[c_0_21, c_0_19])).
cnf(c_0_26, negated_conjecture, (v13_waybel_0(esk2_0,k3_yellow_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_28, plain, ![X22, X23, X24, X25, X26]:((~r1_tarski(X22,X23)|(~r2_hidden(X24,X22)|r2_hidden(X24,X23)))&((r2_hidden(esk5_2(X25,X26),X25)|r1_tarski(X25,X26))&(~r2_hidden(esk5_2(X25,X26),X26)|r1_tarski(X25,X26)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_29, plain, ![X28, X29, X30]:((~v1_xboole_0(X28)|~r2_hidden(X29,X28))&(r2_hidden(esk6_1(X30),X30)|v1_xboole_0(X30))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_30, plain, (r2_hidden(X1,X2)|~r1_tarski(X1,X3)|~r1_tarski(X4,X1)|~r2_hidden(X4,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))|~v13_waybel_0(X2,k2_yellow_1(k1_zfmisc_1(X3)))), inference(rw,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))), inference(rw,[status(thm)],[c_0_25, c_0_24])).
cnf(c_0_32, negated_conjecture, (v13_waybel_0(esk2_0,k2_yellow_1(k1_zfmisc_1(esk1_0)))), inference(rw,[status(thm)],[c_0_26, c_0_19])).
cnf(c_0_33, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_34, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_36, plain, ![X38, X39]:((~m1_subset_1(X38,k1_zfmisc_1(X39))|r1_tarski(X38,X39))&(~r1_tarski(X38,X39)|m1_subset_1(X38,k1_zfmisc_1(X39)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_37, negated_conjecture, (r2_hidden(X1,esk2_0)|~r1_tarski(X1,esk1_0)|~r1_tarski(X2,X1)|~r2_hidden(X2,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])])).
cnf(c_0_38, plain, (r1_tarski(esk5_2(k1_zfmisc_1(X1),X2),X1)|r1_tarski(k1_zfmisc_1(X1),X2)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_39, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_35, c_0_34])).
cnf(c_0_40, negated_conjecture, (v1_xboole_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_41, plain, ![X8, X9]:(((r1_tarski(X8,X9)|X8!=X9)&(r1_tarski(X9,X8)|X8!=X9))&(~r1_tarski(X8,X9)|~r1_tarski(X9,X8)|X8=X9)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_42, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_43, negated_conjecture, (r1_tarski(k1_zfmisc_1(esk1_0),X1)|r2_hidden(esk5_2(k1_zfmisc_1(esk1_0),X1),esk2_0)|~r1_tarski(X2,esk5_2(k1_zfmisc_1(esk1_0),X1))|~r2_hidden(X2,esk2_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_44, negated_conjecture, (r2_hidden(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_45, negated_conjecture, (r1_tarski(esk3_0,X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_46, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_47, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_48, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_49, negated_conjecture, (r1_tarski(esk2_0,k1_zfmisc_1(esk1_0))), inference(spm,[status(thm)],[c_0_42, c_0_31])).
cnf(c_0_50, plain, (r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_51, negated_conjecture, (r1_tarski(k1_zfmisc_1(esk1_0),X1)|r2_hidden(esk5_2(k1_zfmisc_1(esk1_0),X1),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])])).
fof(c_0_52, plain, ![X11, X12]:((~v1_subset_1(X12,X11)|X12!=X11|~m1_subset_1(X12,k1_zfmisc_1(X11)))&(X12=X11|v1_subset_1(X12,X11)|~m1_subset_1(X12,k1_zfmisc_1(X11)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_subset_1])])])])).
cnf(c_0_53, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_54, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_46])).
cnf(c_0_55, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(k2_yellow_1(k1_zfmisc_1(esk1_0))))), inference(rw,[status(thm)],[c_0_47, c_0_19])).
cnf(c_0_56, negated_conjecture, (k1_zfmisc_1(esk1_0)=esk2_0|~r1_tarski(k1_zfmisc_1(esk1_0),esk2_0)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_57, negated_conjecture, (r1_tarski(k1_zfmisc_1(esk1_0),esk2_0)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_58, plain, (~v1_subset_1(X1,X2)|X1!=X2|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_59, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_60, negated_conjecture, (v1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))), inference(rw,[status(thm)],[c_0_55, c_0_24])).
cnf(c_0_61, negated_conjecture, (k1_zfmisc_1(esk1_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_57])])).
cnf(c_0_62, plain, (~v1_subset_1(X1,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_58]), c_0_59])])).
cnf(c_0_63, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_61]), c_0_62]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 64
# Proof object clause steps            : 41
# Proof object formula steps           : 23
# Proof object conjectures             : 22
# Proof object clause conjectures      : 19
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 30
# Processed clauses                    : 2973
# ...of these trivial                  : 8
# ...subsumed                          : 2447
# ...remaining for further processing  : 518
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 47
# Backward-rewritten                   : 55
# Generated clauses                    : 17706
# ...of the previous two non-trivial   : 15107
# Contextual simplify-reflections      : 0
# Paramodulations                      : 17686
# Factorizations                       : 8
# Equation resolutions                 : 12
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 411
#    Positive orientable unit clauses  : 29
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 378
# Current number of unprocessed clauses: 12024
# ...number of literals in the above   : 56704
# Current number of archived formulas  : 0
# Current number of archived clauses   : 105
# Clause-clause subsumption calls (NU) : 21706
# Rec. Clause-clause subsumption calls : 17277
# Non-unit clause-clause subsumptions  : 1532
# Unit Clause-clause subsumption calls : 472
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 28
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 242400

# -------------------------------------------------
# User time                : 0.152 s
# System time              : 0.003 s
# Total time               : 0.155 s
# Maximum resident set size: 28572 pages
