# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d3_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(r1_tarski(X1,X2)<=>![X3, X4]:(r2_hidden(k4_tarski(X3,X4),X1)=>r2_hidden(k4_tarski(X3,X4),X2)))), file('data/MPTP/proofs/377e5770dca44c5380785e15297f9af7.E_input', d3_relat_1)).
fof(d8_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(X3=k5_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>?[X6]:(r2_hidden(k4_tarski(X4,X6),X1)&r2_hidden(k4_tarski(X6,X5),X2))))))), file('data/MPTP/proofs/377e5770dca44c5380785e15297f9af7.E_input', d8_relat_1)).
fof(t50_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>![X4]:(v1_relat_1(X4)=>((r1_tarski(X1,X2)&r1_tarski(X3,X4))=>r1_tarski(k5_relat_1(X1,X3),k5_relat_1(X2,X4))))))), file('data/MPTP/proofs/377e5770dca44c5380785e15297f9af7.E_input', t50_relat_1)).
fof(dt_k5_relat_1, axiom, ![X1, X2]:((v1_relat_1(X1)&v1_relat_1(X2))=>v1_relat_1(k5_relat_1(X1,X2))), file('data/MPTP/proofs/377e5770dca44c5380785e15297f9af7.E_input', dt_k5_relat_1)).
fof(c_0_4, plain, ![X7, X8, X9, X10, X11]:((~r1_tarski(X7,X8)|(~r2_hidden(k4_tarski(X9,X10),X7)|r2_hidden(k4_tarski(X9,X10),X8))|~v1_relat_1(X7))&((r2_hidden(k4_tarski(esk1_2(X7,X11),esk2_2(X7,X11)),X7)|r1_tarski(X7,X11)|~v1_relat_1(X7))&(~r2_hidden(k4_tarski(esk1_2(X7,X11),esk2_2(X7,X11)),X11)|r1_tarski(X7,X11)|~v1_relat_1(X7)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_relat_1])])])])])])).
fof(c_0_5, plain, ![X14, X15, X16, X17, X18, X20, X21, X22, X25]:((((r2_hidden(k4_tarski(X17,esk3_5(X14,X15,X16,X17,X18)),X14)|~r2_hidden(k4_tarski(X17,X18),X16)|X16!=k5_relat_1(X14,X15)|~v1_relat_1(X16)|~v1_relat_1(X15)|~v1_relat_1(X14))&(r2_hidden(k4_tarski(esk3_5(X14,X15,X16,X17,X18),X18),X15)|~r2_hidden(k4_tarski(X17,X18),X16)|X16!=k5_relat_1(X14,X15)|~v1_relat_1(X16)|~v1_relat_1(X15)|~v1_relat_1(X14)))&(~r2_hidden(k4_tarski(X20,X22),X14)|~r2_hidden(k4_tarski(X22,X21),X15)|r2_hidden(k4_tarski(X20,X21),X16)|X16!=k5_relat_1(X14,X15)|~v1_relat_1(X16)|~v1_relat_1(X15)|~v1_relat_1(X14)))&((~r2_hidden(k4_tarski(esk4_3(X14,X15,X16),esk5_3(X14,X15,X16)),X16)|(~r2_hidden(k4_tarski(esk4_3(X14,X15,X16),X25),X14)|~r2_hidden(k4_tarski(X25,esk5_3(X14,X15,X16)),X15))|X16=k5_relat_1(X14,X15)|~v1_relat_1(X16)|~v1_relat_1(X15)|~v1_relat_1(X14))&((r2_hidden(k4_tarski(esk4_3(X14,X15,X16),esk6_3(X14,X15,X16)),X14)|r2_hidden(k4_tarski(esk4_3(X14,X15,X16),esk5_3(X14,X15,X16)),X16)|X16=k5_relat_1(X14,X15)|~v1_relat_1(X16)|~v1_relat_1(X15)|~v1_relat_1(X14))&(r2_hidden(k4_tarski(esk6_3(X14,X15,X16),esk5_3(X14,X15,X16)),X15)|r2_hidden(k4_tarski(esk4_3(X14,X15,X16),esk5_3(X14,X15,X16)),X16)|X16=k5_relat_1(X14,X15)|~v1_relat_1(X16)|~v1_relat_1(X15)|~v1_relat_1(X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_relat_1])])])])])])).
fof(c_0_6, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>![X4]:(v1_relat_1(X4)=>((r1_tarski(X1,X2)&r1_tarski(X3,X4))=>r1_tarski(k5_relat_1(X1,X3),k5_relat_1(X2,X4)))))))), inference(assume_negation,[status(cth)],[t50_relat_1])).
cnf(c_0_7, plain, (r2_hidden(k4_tarski(X3,X4),X2)|~r1_tarski(X1,X2)|~r2_hidden(k4_tarski(X3,X4),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, plain, (r2_hidden(k4_tarski(esk3_5(X1,X2,X3,X4,X5),X5),X2)|~r2_hidden(k4_tarski(X4,X5),X3)|X3!=k5_relat_1(X1,X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_9, negated_conjecture, (v1_relat_1(esk7_0)&(v1_relat_1(esk8_0)&(v1_relat_1(esk9_0)&(v1_relat_1(esk10_0)&((r1_tarski(esk7_0,esk8_0)&r1_tarski(esk9_0,esk10_0))&~r1_tarski(k5_relat_1(esk7_0,esk9_0),k5_relat_1(esk8_0,esk10_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
cnf(c_0_10, plain, (r2_hidden(k4_tarski(esk3_5(X1,X2,X3,X4,X5),X5),X6)|X3!=k5_relat_1(X1,X2)|~r2_hidden(k4_tarski(X4,X5),X3)|~r1_tarski(X2,X6)|~v1_relat_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_11, negated_conjecture, (r1_tarski(esk9_0,esk10_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12, negated_conjecture, (v1_relat_1(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (r2_hidden(k4_tarski(X1,esk3_5(X2,X3,X4,X1,X5)),X2)|~r2_hidden(k4_tarski(X1,X5),X4)|X4!=k5_relat_1(X2,X3)|~v1_relat_1(X4)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_14, plain, (r2_hidden(k4_tarski(X1,X4),X6)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(k4_tarski(X2,X4),X5)|X6!=k5_relat_1(X3,X5)|~v1_relat_1(X6)|~v1_relat_1(X5)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_15, negated_conjecture, (r2_hidden(k4_tarski(esk3_5(X1,esk9_0,X2,X3,X4),X4),esk10_0)|X2!=k5_relat_1(X1,esk9_0)|~r2_hidden(k4_tarski(X3,X4),X2)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])])).
cnf(c_0_16, negated_conjecture, (v1_relat_1(esk10_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, plain, (r2_hidden(k4_tarski(X1,esk3_5(X2,X3,X4,X1,X5)),X6)|X4!=k5_relat_1(X2,X3)|~r2_hidden(k4_tarski(X1,X5),X4)|~r1_tarski(X2,X6)|~v1_relat_1(X2)|~v1_relat_1(X4)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_7, c_0_13])).
cnf(c_0_18, negated_conjecture, (r1_tarski(esk7_0,esk8_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_19, negated_conjecture, (v1_relat_1(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_20, negated_conjecture, (r2_hidden(k4_tarski(X1,X2),X3)|X3!=k5_relat_1(X4,esk10_0)|X5!=k5_relat_1(X6,esk9_0)|~r2_hidden(k4_tarski(X1,esk3_5(X6,esk9_0,X5,X7,X2)),X4)|~r2_hidden(k4_tarski(X7,X2),X5)|~v1_relat_1(X3)|~v1_relat_1(X4)|~v1_relat_1(X5)|~v1_relat_1(X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])])).
cnf(c_0_21, negated_conjecture, (r2_hidden(k4_tarski(X1,esk3_5(esk7_0,X2,X3,X1,X4)),esk8_0)|X3!=k5_relat_1(esk7_0,X2)|~r2_hidden(k4_tarski(X1,X4),X3)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])])).
cnf(c_0_22, negated_conjecture, (v1_relat_1(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_23, negated_conjecture, (r2_hidden(k4_tarski(X1,X2),X3)|X3!=k5_relat_1(esk8_0,esk10_0)|X4!=k5_relat_1(esk7_0,esk9_0)|~r2_hidden(k4_tarski(X1,X2),X4)|~v1_relat_1(X3)|~v1_relat_1(X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_19]), c_0_12])])).
cnf(c_0_24, plain, (r2_hidden(k4_tarski(esk1_2(X1,X2),esk2_2(X1,X2)),X1)|r1_tarski(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_25, plain, (r1_tarski(X1,X2)|~r2_hidden(k4_tarski(esk1_2(X1,X2),esk2_2(X1,X2)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_26, negated_conjecture, (r2_hidden(k4_tarski(esk1_2(X1,X2),esk2_2(X1,X2)),X3)|r1_tarski(X1,X2)|X3!=k5_relat_1(esk8_0,esk10_0)|X1!=k5_relat_1(esk7_0,esk9_0)|~v1_relat_1(X3)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_27, negated_conjecture, (r1_tarski(X1,X2)|X2!=k5_relat_1(esk8_0,esk10_0)|X1!=k5_relat_1(esk7_0,esk9_0)|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_28, negated_conjecture, (r1_tarski(X1,k5_relat_1(esk8_0,esk10_0))|X1!=k5_relat_1(esk7_0,esk9_0)|~v1_relat_1(k5_relat_1(esk8_0,esk10_0))|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_29, negated_conjecture, (~r1_tarski(k5_relat_1(esk7_0,esk9_0),k5_relat_1(esk8_0,esk10_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_30, plain, ![X31, X32]:(~v1_relat_1(X31)|~v1_relat_1(X32)|v1_relat_1(k5_relat_1(X31,X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_relat_1])])).
cnf(c_0_31, negated_conjecture, (~v1_relat_1(k5_relat_1(esk8_0,esk10_0))|~v1_relat_1(k5_relat_1(esk7_0,esk9_0))), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_28]), c_0_29])).
cnf(c_0_32, plain, (v1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_33, negated_conjecture, (~v1_relat_1(k5_relat_1(esk8_0,esk10_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_12]), c_0_19])])).
cnf(c_0_34, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_32]), c_0_16]), c_0_22])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 35
# Proof object clause steps            : 26
# Proof object formula steps           : 9
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 169
# ...of these trivial                  : 1
# ...subsumed                          : 26
# ...remaining for further processing  : 142
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 515
# ...of the previous two non-trivial   : 512
# Contextual simplify-reflections      : 0
# Paramodulations                      : 485
# Factorizations                       : 0
# Equation resolutions                 : 30
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 125
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 117
# Current number of unprocessed clauses: 377
# ...number of literals in the above   : 3430
# Current number of archived formulas  : 0
# Current number of archived clauses   : 17
# Clause-clause subsumption calls (NU) : 6792
# Rec. Clause-clause subsumption calls : 571
# Non-unit clause-clause subsumptions  : 26
# Unit Clause-clause subsumption calls : 102
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 16259

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 28696 pages
