# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/379f301a6a1e4a1cb7cab4376cfff9ab.E_input', t3_subset)).
fof(l1_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/379f301a6a1e4a1cb7cab4376cfff9ab.E_input', l1_zfmisc_1)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/379f301a6a1e4a1cb7cab4376cfff9ab.E_input', d9_lattice3)).
fof(t52_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r1_yellow_0(X1,X4)))&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X3)&![X5]:((v1_finset_1(X5)&m1_subset_1(X5,k1_zfmisc_1(X2)))=>~((r1_yellow_0(X1,X5)&X4=k1_yellow_0(X1,X5)))))))&![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r2_hidden(k1_yellow_0(X1,X4),X3)))))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)<=>r2_lattice3(X1,X3,X4))))))), file('data/MPTP/proofs/379f301a6a1e4a1cb7cab4376cfff9ab.E_input', t52_waybel_0)).
fof(fc1_finset_1, axiom, ![X1]:v1_finset_1(k1_tarski(X1)), file('data/MPTP/proofs/379f301a6a1e4a1cb7cab4376cfff9ab.E_input', fc1_finset_1)).
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/379f301a6a1e4a1cb7cab4376cfff9ab.E_input', d9_yellow_0)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/379f301a6a1e4a1cb7cab4376cfff9ab.E_input', dt_k1_yellow_0)).
fof(t4_yellow_0, axiom, ![X1]:((v4_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)=>![X4]:((r1_lattice3(X1,X4,X3)=>r1_lattice3(X1,X4,X2))&(r2_lattice3(X1,X4,X2)=>r2_lattice3(X1,X4,X3))))))), file('data/MPTP/proofs/379f301a6a1e4a1cb7cab4376cfff9ab.E_input', t4_yellow_0)).
fof(t7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2)))))))), file('data/MPTP/proofs/379f301a6a1e4a1cb7cab4376cfff9ab.E_input', t7_yellow_0)).
fof(fc2_xboole_0, axiom, ![X1]:~(v1_xboole_0(k1_tarski(X1))), file('data/MPTP/proofs/379f301a6a1e4a1cb7cab4376cfff9ab.E_input', fc2_xboole_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/379f301a6a1e4a1cb7cab4376cfff9ab.E_input', fc1_xboole_0)).
fof(t9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(r1_tarski(X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X3,X4)=>r1_lattice3(X1,X2,X4))&(r2_lattice3(X1,X3,X4)=>r2_lattice3(X1,X2,X4)))))), file('data/MPTP/proofs/379f301a6a1e4a1cb7cab4376cfff9ab.E_input', t9_yellow_0)).
fof(c_0_12, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)<=>(![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r1_yellow_0(X1,X4)))&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X3)&![X5]:((v1_finset_1(X5)&m1_subset_1(X5,k1_zfmisc_1(X2)))=>~((r1_yellow_0(X1,X5)&X4=k1_yellow_0(X1,X5)))))))&![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r2_hidden(k1_yellow_0(X1,X4),X3)))))), introduced(definition)).
fof(c_0_13, plain, ![X6, X7]:((~m1_subset_1(X6,k1_zfmisc_1(X7))|r1_tarski(X6,X7))&(~r1_tarski(X6,X7)|m1_subset_1(X6,k1_zfmisc_1(X7)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_14, plain, ![X8, X9]:((~r1_tarski(k1_tarski(X8),X9)|r2_hidden(X8,X9))&(~r2_hidden(X8,X9)|r1_tarski(k1_tarski(X8),X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])).
fof(c_0_15, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)=>(![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r1_yellow_0(X1,X4)))&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X3)&![X5]:((v1_finset_1(X5)&m1_subset_1(X5,k1_zfmisc_1(X2)))=>~((r1_yellow_0(X1,X5)&X4=k1_yellow_0(X1,X5)))))))&![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r2_hidden(k1_yellow_0(X1,X4),X3)))))), inference(split_equiv,[status(thm)],[c_0_12])).
cnf(c_0_16, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, plain, ![X10, X11, X12, X13]:((~r2_lattice3(X10,X11,X12)|(~m1_subset_1(X13,u1_struct_0(X10))|(~r2_hidden(X13,X11)|r1_orders_2(X10,X13,X12)))|~m1_subset_1(X12,u1_struct_0(X10))|~l1_orders_2(X10))&((m1_subset_1(esk1_3(X10,X11,X12),u1_struct_0(X10))|r2_lattice3(X10,X11,X12)|~m1_subset_1(X12,u1_struct_0(X10))|~l1_orders_2(X10))&((r2_hidden(esk1_3(X10,X11,X12),X11)|r2_lattice3(X10,X11,X12)|~m1_subset_1(X12,u1_struct_0(X10))|~l1_orders_2(X10))&(~r1_orders_2(X10,esk1_3(X10,X11,X12),X12)|r2_lattice3(X10,X11,X12)|~m1_subset_1(X12,u1_struct_0(X10))|~l1_orders_2(X10))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
fof(c_0_19, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(epred1_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)<=>r2_lattice3(X1,X3,X4)))))))), inference(apply_def,[status(thm)],[inference(assume_negation,[status(cth)],[t52_waybel_0]), c_0_12])).
fof(c_0_20, plain, ![X39, X40, X41, X42, X43, X45]:((~v1_finset_1(X42)|~m1_subset_1(X42,k1_zfmisc_1(X40))|(X42=k1_xboole_0|r1_yellow_0(X41,X42))|~epred1_3(X41,X40,X39))&((((v1_finset_1(esk7_4(X39,X40,X41,X43))|~r2_hidden(X43,X39)|~m1_subset_1(X43,u1_struct_0(X41))|~epred1_3(X41,X40,X39))&(m1_subset_1(esk7_4(X39,X40,X41,X43),k1_zfmisc_1(X40))|~r2_hidden(X43,X39)|~m1_subset_1(X43,u1_struct_0(X41))|~epred1_3(X41,X40,X39)))&((r1_yellow_0(X41,esk7_4(X39,X40,X41,X43))|~r2_hidden(X43,X39)|~m1_subset_1(X43,u1_struct_0(X41))|~epred1_3(X41,X40,X39))&(X43=k1_yellow_0(X41,esk7_4(X39,X40,X41,X43))|~r2_hidden(X43,X39)|~m1_subset_1(X43,u1_struct_0(X41))|~epred1_3(X41,X40,X39))))&(~v1_finset_1(X45)|~m1_subset_1(X45,k1_zfmisc_1(X40))|(X45=k1_xboole_0|r2_hidden(k1_yellow_0(X41,X45),X39))|~epred1_3(X41,X40,X39)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_15])])])])])])).
cnf(c_0_21, plain, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_22, plain, (r2_hidden(esk1_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_23, plain, ![X19]:v1_finset_1(k1_tarski(X19)), inference(variable_rename,[status(thm)],[fc1_finset_1])).
fof(c_0_24, negated_conjecture, ((~v2_struct_0(esk2_0)&(v3_orders_2(esk2_0)&(v4_orders_2(esk2_0)&l1_orders_2(esk2_0))))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&(epred1_3(esk2_0,esk3_0,esk4_0)&(m1_subset_1(esk5_0,u1_struct_0(esk2_0))&((~r2_lattice3(esk2_0,esk3_0,esk5_0)|~r2_lattice3(esk2_0,esk4_0,esk5_0))&(r2_lattice3(esk2_0,esk3_0,esk5_0)|r2_lattice3(esk2_0,esk4_0,esk5_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_19])])])])).
cnf(c_0_25, plain, (X1=k1_xboole_0|r2_hidden(k1_yellow_0(X3,X1),X4)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~epred1_3(X3,X2,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (r2_lattice3(X1,X2,X3)|m1_subset_1(k1_tarski(esk1_3(X1,X2,X3)),k1_zfmisc_1(X2))|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27, plain, (v1_finset_1(k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28, plain, (r1_orders_2(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_29, negated_conjecture, (r2_lattice3(esk2_0,esk3_0,esk5_0)|r2_lattice3(esk2_0,esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, negated_conjecture, (l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, plain, (k1_tarski(esk1_3(X1,X2,X3))=k1_xboole_0|r2_lattice3(X1,X2,X3)|r2_hidden(k1_yellow_0(X4,k1_tarski(esk1_3(X1,X2,X3))),X5)|~epred1_3(X4,X2,X5)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])])).
cnf(c_0_33, negated_conjecture, (epred1_3(esk2_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_34, plain, ![X20, X21, X22, X23]:(((r2_lattice3(X20,X21,X22)|X22!=k1_yellow_0(X20,X21)|~r1_yellow_0(X20,X21)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))&(~m1_subset_1(X23,u1_struct_0(X20))|(~r2_lattice3(X20,X21,X23)|r1_orders_2(X20,X22,X23))|X22!=k1_yellow_0(X20,X21)|~r1_yellow_0(X20,X21)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20)))&((m1_subset_1(esk6_3(X20,X21,X22),u1_struct_0(X20))|~r2_lattice3(X20,X21,X22)|X22=k1_yellow_0(X20,X21)|~r1_yellow_0(X20,X21)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))&((r2_lattice3(X20,X21,esk6_3(X20,X21,X22))|~r2_lattice3(X20,X21,X22)|X22=k1_yellow_0(X20,X21)|~r1_yellow_0(X20,X21)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))&(~r1_orders_2(X20,X22,esk6_3(X20,X21,X22))|~r2_lattice3(X20,X21,X22)|X22=k1_yellow_0(X20,X21)|~r1_yellow_0(X20,X21)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
fof(c_0_35, plain, ![X25, X26]:(~l1_orders_2(X25)|m1_subset_1(k1_yellow_0(X25,X26),u1_struct_0(X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
cnf(c_0_36, plain, (X1=k1_xboole_0|r1_yellow_0(X3,X1)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~epred1_3(X3,X2,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_37, plain, ![X27, X28, X29, X30]:((~r1_lattice3(X27,X30,X29)|r1_lattice3(X27,X30,X28)|~r1_orders_2(X27,X28,X29)|~m1_subset_1(X29,u1_struct_0(X27))|~m1_subset_1(X28,u1_struct_0(X27))|(~v4_orders_2(X27)|~l1_orders_2(X27)))&(~r2_lattice3(X27,X30,X28)|r2_lattice3(X27,X30,X29)|~r1_orders_2(X27,X28,X29)|~m1_subset_1(X29,u1_struct_0(X27))|~m1_subset_1(X28,u1_struct_0(X27))|(~v4_orders_2(X27)|~l1_orders_2(X27)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_yellow_0])])])])).
cnf(c_0_38, negated_conjecture, (r1_orders_2(esk2_0,X1,esk5_0)|r2_lattice3(esk2_0,esk3_0,esk5_0)|~r2_hidden(X1,esk4_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30]), c_0_31])])).
cnf(c_0_39, negated_conjecture, (k1_tarski(esk1_3(X1,esk3_0,X2))=k1_xboole_0|r2_lattice3(X1,esk3_0,X2)|r2_hidden(k1_yellow_0(esk2_0,k1_tarski(esk1_3(X1,esk3_0,X2))),esk4_0)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_40, plain, (r2_lattice3(X1,X2,X3)|X3!=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_41, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_42, plain, (k1_tarski(esk1_3(X1,X2,X3))=k1_xboole_0|r1_yellow_0(X4,k1_tarski(esk1_3(X1,X2,X3)))|r2_lattice3(X1,X2,X3)|~epred1_3(X4,X2,X5)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_26]), c_0_27])])).
cnf(c_0_43, plain, (r2_lattice3(X1,X2,X4)|~r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,X3,X4)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_44, negated_conjecture, (k1_tarski(esk1_3(X1,esk3_0,X2))=k1_xboole_0|r1_orders_2(esk2_0,k1_yellow_0(esk2_0,k1_tarski(esk1_3(X1,esk3_0,X2))),esk5_0)|r2_lattice3(esk2_0,esk3_0,esk5_0)|r2_lattice3(X1,esk3_0,X2)|~l1_orders_2(X1)|~m1_subset_1(k1_yellow_0(esk2_0,k1_tarski(esk1_3(X1,esk3_0,X2))),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_45, negated_conjecture, (v4_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_46, plain, (r2_lattice3(X1,X2,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_40]), c_0_41])).
cnf(c_0_47, negated_conjecture, (k1_tarski(esk1_3(X1,esk3_0,X2))=k1_xboole_0|r1_yellow_0(esk2_0,k1_tarski(esk1_3(X1,esk3_0,X2)))|r2_lattice3(X1,esk3_0,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_42, c_0_33])).
cnf(c_0_48, negated_conjecture, (k1_tarski(esk1_3(X1,esk3_0,X2))=k1_xboole_0|r2_lattice3(esk2_0,esk3_0,esk5_0)|r2_lattice3(esk2_0,X3,esk5_0)|r2_lattice3(X1,esk3_0,X2)|~r2_lattice3(esk2_0,X3,k1_yellow_0(esk2_0,k1_tarski(esk1_3(X1,esk3_0,X2))))|~l1_orders_2(X1)|~m1_subset_1(k1_yellow_0(esk2_0,k1_tarski(esk1_3(X1,esk3_0,X2))),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_30]), c_0_31])])).
cnf(c_0_49, negated_conjecture, (k1_tarski(esk1_3(X1,esk3_0,X2))=k1_xboole_0|r2_lattice3(esk2_0,k1_tarski(esk1_3(X1,esk3_0,X2)),k1_yellow_0(esk2_0,k1_tarski(esk1_3(X1,esk3_0,X2))))|r2_lattice3(X1,esk3_0,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_30])])).
fof(c_0_50, plain, ![X31, X32, X33]:((~r1_lattice3(X31,k1_tarski(X33),X32)|r1_orders_2(X31,X32,X33)|~m1_subset_1(X33,u1_struct_0(X31))|~m1_subset_1(X32,u1_struct_0(X31))|~l1_orders_2(X31))&((~r1_orders_2(X31,X32,X33)|r1_lattice3(X31,k1_tarski(X33),X32)|~m1_subset_1(X33,u1_struct_0(X31))|~m1_subset_1(X32,u1_struct_0(X31))|~l1_orders_2(X31))&((~r2_lattice3(X31,k1_tarski(X33),X32)|r1_orders_2(X31,X33,X32)|~m1_subset_1(X33,u1_struct_0(X31))|~m1_subset_1(X32,u1_struct_0(X31))|~l1_orders_2(X31))&(~r1_orders_2(X31,X33,X32)|r2_lattice3(X31,k1_tarski(X33),X32)|~m1_subset_1(X33,u1_struct_0(X31))|~m1_subset_1(X32,u1_struct_0(X31))|~l1_orders_2(X31))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])).
cnf(c_0_51, negated_conjecture, (k1_tarski(esk1_3(X1,esk3_0,X2))=k1_xboole_0|r2_lattice3(esk2_0,k1_tarski(esk1_3(X1,esk3_0,X2)),esk5_0)|r2_lattice3(esk2_0,esk3_0,esk5_0)|r2_lattice3(X1,esk3_0,X2)|~l1_orders_2(X1)|~m1_subset_1(k1_yellow_0(esk2_0,k1_tarski(esk1_3(X1,esk3_0,X2))),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
fof(c_0_52, plain, ![X34]:~v1_xboole_0(k1_tarski(X34)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_xboole_0])])).
cnf(c_0_53, plain, (r1_orders_2(X1,X2,X3)|~r2_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_54, negated_conjecture, (k1_tarski(esk1_3(X1,esk3_0,X2))=k1_xboole_0|r2_lattice3(esk2_0,k1_tarski(esk1_3(X1,esk3_0,X2)),esk5_0)|r2_lattice3(esk2_0,esk3_0,esk5_0)|r2_lattice3(X1,esk3_0,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_41]), c_0_30])])).
cnf(c_0_55, plain, (~v1_xboole_0(k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_56, negated_conjecture, (k1_tarski(esk1_3(X1,esk3_0,X2))=k1_xboole_0|r1_orders_2(esk2_0,esk1_3(X1,esk3_0,X2),esk5_0)|r2_lattice3(esk2_0,esk3_0,esk5_0)|r2_lattice3(X1,esk3_0,X2)|~l1_orders_2(X1)|~m1_subset_1(esk1_3(X1,esk3_0,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_30]), c_0_31])])).
cnf(c_0_57, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_58, plain, (r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk1_3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_59, negated_conjecture, (r1_orders_2(esk2_0,esk1_3(X1,esk3_0,X2),esk5_0)|r2_lattice3(esk2_0,esk3_0,esk5_0)|r2_lattice3(X1,esk3_0,X2)|~l1_orders_2(X1)|~m1_subset_1(esk1_3(X1,esk3_0,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57])])).
fof(c_0_60, plain, ![X35, X36, X37, X38]:((~r1_lattice3(X35,X37,X38)|r1_lattice3(X35,X36,X38)|~m1_subset_1(X38,u1_struct_0(X35))|~r1_tarski(X36,X37)|~l1_orders_2(X35))&(~r2_lattice3(X35,X37,X38)|r2_lattice3(X35,X36,X38)|~m1_subset_1(X38,u1_struct_0(X35))|~r1_tarski(X36,X37)|~l1_orders_2(X35))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_yellow_0])])])])).
cnf(c_0_61, negated_conjecture, (r2_lattice3(esk2_0,esk3_0,esk5_0)|~m1_subset_1(esk1_3(esk2_0,esk3_0,esk5_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_30]), c_0_31])])).
cnf(c_0_62, plain, (m1_subset_1(esk1_3(X1,X2,X3),u1_struct_0(X1))|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_63, plain, (r2_lattice3(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_tarski(X4,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_64, negated_conjecture, (r2_lattice3(esk2_0,esk3_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_30]), c_0_31])])).
cnf(c_0_65, negated_conjecture, (r2_lattice3(esk2_0,X1,esk5_0)|~r1_tarski(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_30]), c_0_31])])).
cnf(c_0_66, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_67, negated_conjecture, (r2_lattice3(esk2_0,X1,esk5_0)|~m1_subset_1(X1,k1_zfmisc_1(esk3_0))), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_68, plain, (m1_subset_1(esk7_4(X1,X2,X3,X4),k1_zfmisc_1(X2))|~r2_hidden(X4,X1)|~m1_subset_1(X4,u1_struct_0(X3))|~epred1_3(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_69, plain, (r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~r1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_70, plain, (r2_lattice3(esk2_0,esk7_4(X1,esk3_0,X2,X3),esk5_0)|~epred1_3(X2,esk3_0,X1)|~r2_hidden(X3,X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_71, plain, (r1_orders_2(esk2_0,X1,esk5_0)|X1!=k1_yellow_0(esk2_0,esk7_4(X2,esk3_0,X3,X4))|~epred1_3(X3,esk3_0,X2)|~r1_yellow_0(esk2_0,esk7_4(X2,esk3_0,X3,X4))|~r2_hidden(X4,X2)|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(X3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_30]), c_0_31])])).
cnf(c_0_72, plain, (X1=k1_yellow_0(X2,esk7_4(X3,X4,X2,X1))|~r2_hidden(X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_73, plain, (r1_yellow_0(X1,esk7_4(X2,X3,X1,X4))|~r2_hidden(X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~epred1_3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_74, plain, (r1_orders_2(esk2_0,X1,esk5_0)|~epred1_3(esk2_0,esk3_0,X2)|~r2_hidden(X1,X2)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72])]), c_0_73])).
cnf(c_0_75, negated_conjecture, (r1_orders_2(esk2_0,X1,esk5_0)|~r2_hidden(X1,esk4_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_74, c_0_33])).
cnf(c_0_76, negated_conjecture, (~r2_lattice3(esk2_0,esk3_0,esk5_0)|~r2_lattice3(esk2_0,esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_77, negated_conjecture, (r1_orders_2(esk2_0,esk1_3(X1,esk4_0,X2),esk5_0)|r2_lattice3(X1,esk4_0,X2)|~l1_orders_2(X1)|~m1_subset_1(esk1_3(X1,esk4_0,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_75, c_0_22])).
cnf(c_0_78, negated_conjecture, (~r2_lattice3(esk2_0,esk4_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_64])])).
cnf(c_0_79, negated_conjecture, (~m1_subset_1(esk1_3(esk2_0,esk4_0,esk5_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_77]), c_0_30]), c_0_31])]), c_0_78])).
cnf(c_0_80, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_62]), c_0_30]), c_0_31])]), c_0_78]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 81
# Proof object clause steps            : 54
# Proof object formula steps           : 27
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 26
# Proof object simplifying inferences  : 45
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 41
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 41
# Processed clauses                    : 683
# ...of these trivial                  : 2
# ...subsumed                          : 113
# ...remaining for further processing  : 568
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 20
# Backward-rewritten                   : 181
# Generated clauses                    : 2113
# ...of the previous two non-trivial   : 2080
# Contextual simplify-reflections      : 46
# Paramodulations                      : 2088
# Factorizations                       : 0
# Equation resolutions                 : 25
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 326
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 312
# Current number of unprocessed clauses: 1472
# ...number of literals in the above   : 17531
# Current number of archived formulas  : 0
# Current number of archived clauses   : 242
# Clause-clause subsumption calls (NU) : 90599
# Rec. Clause-clause subsumption calls : 2383
# Non-unit clause-clause subsumptions  : 179
# Unit Clause-clause subsumption calls : 176
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 111415

# -------------------------------------------------
# User time                : 0.113 s
# System time              : 0.013 s
# Total time               : 0.126 s
# Maximum resident set size: 28832 pages
