# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2m
# and selection function SelectCQArNTNpEqFirst.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t31_xboole_1, axiom, ![X1, X2, X3]:r1_tarski(k2_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3)),k2_xboole_0(X2,X3)), file('data/MPTP/proofs/37b11d79a4ac412d859c78638178a508.E_input', t31_xboole_1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/37b11d79a4ac412d859c78638178a508.E_input', l51_zfmisc_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/37b11d79a4ac412d859c78638178a508.E_input', t48_xboole_1)).
fof(t23_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3)), file('data/MPTP/proofs/37b11d79a4ac412d859c78638178a508.E_input', t23_xboole_1)).
fof(t22_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/37b11d79a4ac412d859c78638178a508.E_input', t22_xboole_1)).
fof(t19_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X1,X3))=>r1_tarski(X1,k3_xboole_0(X2,X3))), file('data/MPTP/proofs/37b11d79a4ac412d859c78638178a508.E_input', t19_xboole_1)).
fof(t95_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>r1_tarski(k3_tarski(X1),k3_tarski(X2))), file('data/MPTP/proofs/37b11d79a4ac412d859c78638178a508.E_input', t95_zfmisc_1)).
fof(t97_zfmisc_1, conjecture, ![X1, X2]:r1_tarski(k3_tarski(k3_xboole_0(X1,X2)),k3_xboole_0(k3_tarski(X1),k3_tarski(X2))), file('data/MPTP/proofs/37b11d79a4ac412d859c78638178a508.E_input', t97_zfmisc_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/37b11d79a4ac412d859c78638178a508.E_input', t36_xboole_1)).
fof(c_0_9, plain, ![X4, X5, X6]:r1_tarski(k2_xboole_0(k3_xboole_0(X4,X5),k3_xboole_0(X4,X6)),k2_xboole_0(X5,X6)), inference(variable_rename,[status(thm)],[t31_xboole_1])).
fof(c_0_10, plain, ![X7, X8]:k3_tarski(k2_tarski(X7,X8))=k2_xboole_0(X7,X8), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_11, plain, ![X9, X10]:k4_xboole_0(X9,k4_xboole_0(X9,X10))=k3_xboole_0(X9,X10), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_12, plain, ![X11, X12, X13]:k3_xboole_0(X11,k2_xboole_0(X12,X13))=k2_xboole_0(k3_xboole_0(X11,X12),k3_xboole_0(X11,X13)), inference(variable_rename,[status(thm)],[t23_xboole_1])).
cnf(c_0_13, plain, (r1_tarski(k2_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3)),k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (k3_xboole_0(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_17, plain, ![X14, X15]:k2_xboole_0(X14,k3_xboole_0(X14,X15))=X14, inference(variable_rename,[status(thm)],[t22_xboole_1])).
cnf(c_0_18, plain, (r1_tarski(k3_tarski(k2_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k4_xboole_0(X1,k4_xboole_0(X1,X3)))),k3_tarski(k2_tarski(X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_14]), c_0_14]), c_0_15]), c_0_15])).
cnf(c_0_19, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k3_tarski(k2_tarski(X2,X3))))=k3_tarski(k2_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k4_xboole_0(X1,k4_xboole_0(X1,X3))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_14]), c_0_14]), c_0_15]), c_0_15]), c_0_15])).
cnf(c_0_20, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_21, plain, ![X16, X17, X18]:(~r1_tarski(X16,X17)|~r1_tarski(X16,X18)|r1_tarski(X16,k3_xboole_0(X17,X18))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t19_xboole_1])])).
fof(c_0_22, plain, ![X19, X20]:(~r1_tarski(X19,X20)|r1_tarski(k3_tarski(X19),k3_tarski(X20))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t95_zfmisc_1])])).
cnf(c_0_23, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,k3_tarski(k2_tarski(X2,X3)))),k3_tarski(k2_tarski(X2,X3)))), inference(rw,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_24, plain, (k3_tarski(k2_tarski(X1,k4_xboole_0(X1,k4_xboole_0(X1,X2))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_14]), c_0_15])).
fof(c_0_25, negated_conjecture, ~(![X1, X2]:r1_tarski(k3_tarski(k3_xboole_0(X1,X2)),k3_xboole_0(k3_tarski(X1),k3_tarski(X2)))), inference(assume_negation,[status(cth)],[t97_zfmisc_1])).
cnf(c_0_26, plain, (r1_tarski(X1,k3_xboole_0(X2,X3))|~r1_tarski(X1,X2)|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (r1_tarski(k3_tarski(X1),k3_tarski(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_29, plain, ![X23, X24]:r1_tarski(k4_xboole_0(X23,X24),X23), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_30, negated_conjecture, ~r1_tarski(k3_tarski(k3_xboole_0(esk1_0,esk2_0)),k3_xboole_0(k3_tarski(esk1_0),k3_tarski(esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])).
cnf(c_0_31, plain, (r1_tarski(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))|~r1_tarski(X1,X3)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_26, c_0_15])).
cnf(c_0_32, plain, (r1_tarski(k3_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2))),k3_tarski(X2))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_33, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34, negated_conjecture, (~r1_tarski(k3_tarski(k3_xboole_0(esk1_0,esk2_0)),k3_xboole_0(k3_tarski(esk1_0),k3_tarski(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_35, plain, (r1_tarski(k3_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2))),k4_xboole_0(X3,k4_xboole_0(X3,k3_tarski(X2))))|~r1_tarski(k3_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2))),X3)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_36, plain, (r1_tarski(k3_tarski(k4_xboole_0(X1,X2)),k3_tarski(X1))), inference(spm,[status(thm)],[c_0_27, c_0_33])).
cnf(c_0_37, negated_conjecture, (~r1_tarski(k3_tarski(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0))),k4_xboole_0(k3_tarski(esk1_0),k4_xboole_0(k3_tarski(esk1_0),k3_tarski(esk2_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_15]), c_0_15])).
cnf(c_0_38, plain, (r1_tarski(k3_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2))),k4_xboole_0(k3_tarski(X1),k4_xboole_0(k3_tarski(X1),k3_tarski(X2))))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_39, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_38])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 40
# Proof object clause steps            : 21
# Proof object formula steps           : 19
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 9
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 7
# Processed clauses                    : 73
# ...of these trivial                  : 1
# ...subsumed                          : 0
# ...remaining for further processing  : 72
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 281
# ...of the previous two non-trivial   : 270
# Contextual simplify-reflections      : 0
# Paramodulations                      : 281
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 63
#    Positive orientable unit clauses  : 49
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 14
# Current number of unprocessed clauses: 210
# ...number of literals in the above   : 245
# Current number of archived formulas  : 0
# Current number of archived clauses   : 11
# Clause-clause subsumption calls (NU) : 36
# Rec. Clause-clause subsumption calls : 36
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 22
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 392
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6998

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.004 s
# Total time               : 0.011 s
# Maximum resident set size: 28844 pages
