# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t53_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>(![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>r1_yellow_0(X1,X2))=>v3_lattice3(X1))), file('data/MPTP/proofs/37d6f1277f574c29b7dc2de993d326de.E_input', t53_yellow_0)).
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/37d6f1277f574c29b7dc2de993d326de.E_input', d9_yellow_0)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/37d6f1277f574c29b7dc2de993d326de.E_input', dt_k1_yellow_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/37d6f1277f574c29b7dc2de993d326de.E_input', t3_subset)).
fof(d12_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v3_lattice3(X1)<=>![X2]:?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4))))))), file('data/MPTP/proofs/37d6f1277f574c29b7dc2de993d326de.E_input', d12_lattice3)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/37d6f1277f574c29b7dc2de993d326de.E_input', t17_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/37d6f1277f574c29b7dc2de993d326de.E_input', commutativity_k3_xboole_0)).
fof(t50_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((r1_yellow_0(X1,X2)=>r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))&((r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))=>r1_yellow_0(X1,X2))&((r2_yellow_0(X1,X2)=>r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))&(r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))=>r2_yellow_0(X1,X2)))))), file('data/MPTP/proofs/37d6f1277f574c29b7dc2de993d326de.E_input', t50_yellow_0)).
fof(c_0_8, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>(![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>r1_yellow_0(X1,X2))=>v3_lattice3(X1)))), inference(assume_negation,[status(cth)],[t53_yellow_0])).
fof(c_0_9, plain, ![X7, X8, X9, X10]:(((r2_lattice3(X7,X8,X9)|X9!=k1_yellow_0(X7,X8)|~r1_yellow_0(X7,X8)|~m1_subset_1(X9,u1_struct_0(X7))|~l1_orders_2(X7))&(~m1_subset_1(X10,u1_struct_0(X7))|(~r2_lattice3(X7,X8,X10)|r1_orders_2(X7,X9,X10))|X9!=k1_yellow_0(X7,X8)|~r1_yellow_0(X7,X8)|~m1_subset_1(X9,u1_struct_0(X7))|~l1_orders_2(X7)))&((m1_subset_1(esk2_3(X7,X8,X9),u1_struct_0(X7))|~r2_lattice3(X7,X8,X9)|X9=k1_yellow_0(X7,X8)|~r1_yellow_0(X7,X8)|~m1_subset_1(X9,u1_struct_0(X7))|~l1_orders_2(X7))&((r2_lattice3(X7,X8,esk2_3(X7,X8,X9))|~r2_lattice3(X7,X8,X9)|X9=k1_yellow_0(X7,X8)|~r1_yellow_0(X7,X8)|~m1_subset_1(X9,u1_struct_0(X7))|~l1_orders_2(X7))&(~r1_orders_2(X7,X9,esk2_3(X7,X8,X9))|~r2_lattice3(X7,X8,X9)|X9=k1_yellow_0(X7,X8)|~r1_yellow_0(X7,X8)|~m1_subset_1(X9,u1_struct_0(X7))|~l1_orders_2(X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
fof(c_0_10, plain, ![X12, X13]:(~l1_orders_2(X12)|m1_subset_1(k1_yellow_0(X12,X13),u1_struct_0(X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
fof(c_0_11, negated_conjecture, ![X6]:((~v2_struct_0(esk1_0)&l1_orders_2(esk1_0))&((~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_yellow_0(esk1_0,X6))&~v3_lattice3(esk1_0))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])).
fof(c_0_12, plain, ![X14, X15]:((~m1_subset_1(X14,k1_zfmisc_1(X15))|r1_tarski(X14,X15))&(~r1_tarski(X14,X15)|m1_subset_1(X14,k1_zfmisc_1(X15)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_13, plain, ![X16, X17, X19, X21]:(((m1_subset_1(esk3_2(X16,X17),u1_struct_0(X16))|~v3_lattice3(X16)|~l1_orders_2(X16))&((r2_lattice3(X16,X17,esk3_2(X16,X17))|~v3_lattice3(X16)|~l1_orders_2(X16))&(~m1_subset_1(X19,u1_struct_0(X16))|(~r2_lattice3(X16,X17,X19)|r1_orders_2(X16,esk3_2(X16,X17),X19))|~v3_lattice3(X16)|~l1_orders_2(X16))))&((m1_subset_1(esk5_2(X16,X21),u1_struct_0(X16))|~r2_lattice3(X16,esk4_1(X16),X21)|~m1_subset_1(X21,u1_struct_0(X16))|v3_lattice3(X16)|~l1_orders_2(X16))&((r2_lattice3(X16,esk4_1(X16),esk5_2(X16,X21))|~r2_lattice3(X16,esk4_1(X16),X21)|~m1_subset_1(X21,u1_struct_0(X16))|v3_lattice3(X16)|~l1_orders_2(X16))&(~r1_orders_2(X16,X21,esk5_2(X16,X21))|~r2_lattice3(X16,esk4_1(X16),X21)|~m1_subset_1(X21,u1_struct_0(X16))|v3_lattice3(X16)|~l1_orders_2(X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_lattice3])])])])])).
cnf(c_0_14, plain, (r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~r1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, negated_conjecture, (r1_yellow_0(esk1_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_18, plain, ![X23, X24]:r1_tarski(k3_xboole_0(X23,X24),X23), inference(variable_rename,[status(thm)],[t17_xboole_1])).
cnf(c_0_19, plain, (v3_lattice3(X1)|~r1_orders_2(X1,X2,esk5_2(X1,X2))|~r2_lattice3(X1,esk4_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),X3)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_14]), c_0_15])).
cnf(c_0_21, plain, (r2_lattice3(X1,X2,X3)|X3!=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_22, negated_conjecture, (r1_yellow_0(esk1_0,X1)|~r1_tarski(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_23, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_24, plain, ![X25, X26]:k3_xboole_0(X25,X26)=k3_xboole_0(X26,X25), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_25, plain, (v3_lattice3(X1)|~r2_lattice3(X1,X2,esk5_2(X1,k1_yellow_0(X1,X2)))|~r2_lattice3(X1,esk4_1(X1),k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X2)|~m1_subset_1(esk5_2(X1,k1_yellow_0(X1,X2)),u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_15])).
cnf(c_0_26, plain, (r2_lattice3(X1,esk4_1(X1),esk5_2(X1,X2))|v3_lattice3(X1)|~r2_lattice3(X1,esk4_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_27, plain, (r2_lattice3(X1,X2,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_21]), c_0_15])).
fof(c_0_28, plain, ![X27, X28]:((~r1_yellow_0(X27,X28)|r1_yellow_0(X27,k3_xboole_0(X28,u1_struct_0(X27)))|(v2_struct_0(X27)|~l1_orders_2(X27)))&((~r1_yellow_0(X27,k3_xboole_0(X28,u1_struct_0(X27)))|r1_yellow_0(X27,X28)|(v2_struct_0(X27)|~l1_orders_2(X27)))&((~r2_yellow_0(X27,X28)|r2_yellow_0(X27,k3_xboole_0(X28,u1_struct_0(X27)))|(v2_struct_0(X27)|~l1_orders_2(X27)))&(~r2_yellow_0(X27,k3_xboole_0(X28,u1_struct_0(X27)))|r2_yellow_0(X27,X28)|(v2_struct_0(X27)|~l1_orders_2(X27)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_yellow_0])])])])])).
cnf(c_0_29, negated_conjecture, (r1_yellow_0(esk1_0,k3_xboole_0(u1_struct_0(esk1_0),X1))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_30, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, plain, (v3_lattice3(X1)|~r1_yellow_0(X1,esk4_1(X1))|~m1_subset_1(esk5_2(X1,k1_yellow_0(X1,esk4_1(X1))),u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_15]), c_0_27])).
cnf(c_0_32, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|v3_lattice3(X1)|~r2_lattice3(X1,esk4_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_33, plain, (r1_yellow_0(X1,X2)|v2_struct_0(X1)|~r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34, negated_conjecture, (r1_yellow_0(esk1_0,k3_xboole_0(X1,u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_35, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_36, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_37, plain, (v3_lattice3(X1)|~r1_yellow_0(X1,esk4_1(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_15]), c_0_27])).
cnf(c_0_38, negated_conjecture, (r1_yellow_0(esk1_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])]), c_0_36])).
cnf(c_0_39, negated_conjecture, (~v3_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_40, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_35])]), c_0_39]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 41
# Proof object clause steps            : 24
# Proof object formula steps           : 17
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 15
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 24
# Processed clauses                    : 76
# ...of these trivial                  : 2
# ...subsumed                          : 12
# ...remaining for further processing  : 62
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 4
# Generated clauses                    : 49
# ...of the previous two non-trivial   : 31
# Contextual simplify-reflections      : 7
# Paramodulations                      : 47
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 31
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 24
# Current number of unprocessed clauses: 2
# ...number of literals in the above   : 16
# Current number of archived formulas  : 0
# Current number of archived clauses   : 29
# Clause-clause subsumption calls (NU) : 309
# Rec. Clause-clause subsumption calls : 66
# Non-unit clause-clause subsumptions  : 20
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2826

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28696 pages
