# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t11_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k2_xboole_0(X1,X2),X3)=>r1_tarski(X1,X3)), file('data/MPTP/proofs/37e78b554939474cb752df7ee4fb2ec4.E_input', t11_xboole_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/37e78b554939474cb752df7ee4fb2ec4.E_input', t12_xboole_1)).
fof(t10_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(X1,k2_xboole_0(X3,X2))), file('data/MPTP/proofs/37e78b554939474cb752df7ee4fb2ec4.E_input', t10_xboole_1)).
fof(t146_funct_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(X1,k1_relat_1(X2))=>r1_tarski(X1,k10_relat_1(X2,k9_relat_1(X2,X1))))), file('data/MPTP/proofs/37e78b554939474cb752df7ee4fb2ec4.E_input', t146_funct_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/37e78b554939474cb752df7ee4fb2ec4.E_input', d10_xboole_0)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/37e78b554939474cb752df7ee4fb2ec4.E_input', t7_xboole_1)).
fof(t8_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X3,X2))=>r1_tarski(k2_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/37e78b554939474cb752df7ee4fb2ec4.E_input', t8_xboole_1)).
fof(t167_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k10_relat_1(X2,X1),k1_relat_1(X2))), file('data/MPTP/proofs/37e78b554939474cb752df7ee4fb2ec4.E_input', t167_relat_1)).
fof(t170_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k10_relat_1(X2,X1),k10_relat_1(X2,k2_relat_1(X2)))), file('data/MPTP/proofs/37e78b554939474cb752df7ee4fb2ec4.E_input', t170_relat_1)).
fof(t169_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k10_relat_1(X1,k2_relat_1(X1))=k1_relat_1(X1)), file('data/MPTP/proofs/37e78b554939474cb752df7ee4fb2ec4.E_input', t169_relat_1)).
fof(t139_funct_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>k10_relat_1(k7_relat_1(X3,X1),X2)=k3_xboole_0(X1,k10_relat_1(X3,X2))), file('data/MPTP/proofs/37e78b554939474cb752df7ee4fb2ec4.E_input', t139_funct_1)).
fof(t97_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(k1_relat_1(X2),X1)=>k7_relat_1(X2,X1)=X2)), file('data/MPTP/proofs/37e78b554939474cb752df7ee4fb2ec4.E_input', t97_relat_1)).
fof(t192_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k7_relat_1(X2,X1)=k7_relat_1(X2,k3_xboole_0(k1_relat_1(X2),X1))), file('data/MPTP/proofs/37e78b554939474cb752df7ee4fb2ec4.E_input', t192_relat_1)).
fof(t98_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k7_relat_1(X1,k1_relat_1(X1))=X1), file('data/MPTP/proofs/37e78b554939474cb752df7ee4fb2ec4.E_input', t98_relat_1)).
fof(t90_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k1_relat_1(k7_relat_1(X2,X1))=k3_xboole_0(k1_relat_1(X2),X1)), file('data/MPTP/proofs/37e78b554939474cb752df7ee4fb2ec4.E_input', t90_relat_1)).
fof(t91_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(X1,k1_relat_1(X2))=>k1_relat_1(k7_relat_1(X2,X1))=X1)), file('data/MPTP/proofs/37e78b554939474cb752df7ee4fb2ec4.E_input', t91_relat_1)).
fof(dt_k7_relat_1, axiom, ![X1, X2]:(v1_relat_1(X1)=>v1_relat_1(k7_relat_1(X1,X2))), file('data/MPTP/proofs/37e78b554939474cb752df7ee4fb2ec4.E_input', dt_k7_relat_1)).
fof(t87_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k1_relat_1(k7_relat_1(X2,X1)),X1)), file('data/MPTP/proofs/37e78b554939474cb752df7ee4fb2ec4.E_input', t87_relat_1)).
fof(t148_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k2_relat_1(k7_relat_1(X2,X1))=k9_relat_1(X2,X1)), file('data/MPTP/proofs/37e78b554939474cb752df7ee4fb2ec4.E_input', t148_relat_1)).
fof(c_0_19, plain, ![X42, X43, X44]:(~r1_tarski(k2_xboole_0(X42,X43),X44)|r1_tarski(X42,X44)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_xboole_1])])).
fof(c_0_20, plain, ![X6, X7]:(~r1_tarski(X6,X7)|k2_xboole_0(X6,X7)=X7), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_21, plain, (r1_tarski(X1,X3)|~r1_tarski(k2_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_22, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_23, plain, ![X34, X35, X36]:(~r1_tarski(X34,X35)|r1_tarski(X34,k2_xboole_0(X36,X35))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_xboole_1])])).
fof(c_0_24, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(X1,k1_relat_1(X2))=>r1_tarski(X1,k10_relat_1(X2,k9_relat_1(X2,X1)))))), inference(assume_negation,[status(cth)],[t146_funct_1])).
fof(c_0_25, plain, ![X49, X50]:(((r1_tarski(X49,X50)|X49!=X50)&(r1_tarski(X50,X49)|X49!=X50))&(~r1_tarski(X49,X50)|~r1_tarski(X50,X49)|X49=X50)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_26, plain, ![X32, X33]:r1_tarski(X32,k2_xboole_0(X32,X33)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_27, plain, (r1_tarski(X1,X2)|~r1_tarski(X3,X2)|~r1_tarski(X1,X3)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28, plain, (r1_tarski(X1,k2_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_29, negated_conjecture, (v1_relat_1(esk2_0)&(r1_tarski(esk1_0,k1_relat_1(esk2_0))&~r1_tarski(esk1_0,k10_relat_1(esk2_0,k9_relat_1(esk2_0,esk1_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).
cnf(c_0_30, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_32, plain, ![X51, X52, X53]:(~r1_tarski(X51,X52)|~r1_tarski(X53,X52)|r1_tarski(k2_xboole_0(X51,X53),X52)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_xboole_1])])).
cnf(c_0_33, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,X4)|~r1_tarski(X4,X3)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_35, negated_conjecture, (r1_tarski(esk1_0,k1_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, plain, (k2_xboole_0(X1,X2)=X1|~r1_tarski(k2_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_37, plain, (r1_tarski(k2_xboole_0(X1,X3),X2)|~r1_tarski(X1,X2)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_33])).
fof(c_0_39, plain, ![X11, X12]:(~v1_relat_1(X12)|r1_tarski(k10_relat_1(X12,X11),k1_relat_1(X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t167_relat_1])])).
cnf(c_0_40, negated_conjecture, (r1_tarski(esk1_0,k2_xboole_0(X1,X2))|~r1_tarski(k1_relat_1(esk2_0),X2)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_41, plain, (k2_xboole_0(X1,X2)=X1|~r1_tarski(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])])).
fof(c_0_42, plain, ![X28, X29]:(~v1_relat_1(X29)|r1_tarski(k10_relat_1(X29,X28),k10_relat_1(X29,k2_relat_1(X29)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t170_relat_1])])).
fof(c_0_43, plain, ![X18]:(~v1_relat_1(X18)|k10_relat_1(X18,k2_relat_1(X18))=k1_relat_1(X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t169_relat_1])])).
cnf(c_0_44, plain, (k2_xboole_0(X1,X2)=X3|~r1_tarski(k2_xboole_0(X1,X2),X3)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_30, c_0_28])).
cnf(c_0_45, plain, (r1_tarski(k10_relat_1(X1,X2),k1_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_46, negated_conjecture, (r1_tarski(esk1_0,X1)|~r1_tarski(k1_relat_1(esk2_0),X2)|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_47, plain, (r1_tarski(k10_relat_1(X1,X2),k10_relat_1(X1,k2_relat_1(X1)))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_48, plain, (k10_relat_1(X1,k2_relat_1(X1))=k1_relat_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)|~r1_tarski(X3,X1)), inference(spm,[status(thm)],[c_0_44, c_0_41])).
cnf(c_0_50, plain, (r1_tarski(X1,k1_relat_1(X2))|~r1_tarski(X1,k10_relat_1(X2,X3))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_27, c_0_45])).
cnf(c_0_51, negated_conjecture, (r1_tarski(esk1_0,X1)|~r1_tarski(k1_relat_1(esk2_0),X1)), inference(spm,[status(thm)],[c_0_46, c_0_38])).
cnf(c_0_52, plain, (r1_tarski(k1_relat_1(X1),k10_relat_1(X1,k2_relat_1(X1)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_53, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_54, plain, (k2_xboole_0(X1,X2)=X1|~r1_tarski(k2_xboole_0(X1,X2),X3)|~r1_tarski(X3,X1)), inference(spm,[status(thm)],[c_0_49, c_0_31])).
cnf(c_0_55, plain, (r1_tarski(k2_xboole_0(X1,X2),k1_relat_1(X3))|~r1_tarski(X2,k10_relat_1(X3,X4))|~r1_tarski(X1,k10_relat_1(X3,X4))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_50, c_0_37])).
cnf(c_0_56, negated_conjecture, (r1_tarski(esk1_0,k10_relat_1(esk2_0,k2_relat_1(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53])])).
cnf(c_0_57, plain, (k2_xboole_0(X1,X2)=X1|~r1_tarski(k2_xboole_0(k2_xboole_0(X1,X2),X3),X1)), inference(spm,[status(thm)],[c_0_54, c_0_31])).
cnf(c_0_58, negated_conjecture, (r1_tarski(k2_xboole_0(X1,esk1_0),k1_relat_1(esk2_0))|~r1_tarski(X1,k10_relat_1(esk2_0,k2_relat_1(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_53])])).
cnf(c_0_59, plain, (X1=X2|~r1_tarski(k2_xboole_0(X1,X3),X2)|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_57, c_0_22])).
cnf(c_0_60, negated_conjecture, (r1_tarski(k2_xboole_0(k10_relat_1(esk2_0,X1),esk1_0),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_47]), c_0_53])])).
fof(c_0_61, plain, ![X15, X16, X17]:(~v1_relat_1(X17)|k10_relat_1(k7_relat_1(X17,X15),X16)=k3_xboole_0(X15,k10_relat_1(X17,X16))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t139_funct_1])])).
cnf(c_0_62, negated_conjecture, (k10_relat_1(esk2_0,X1)=k1_relat_1(esk2_0)|~r1_tarski(k1_relat_1(esk2_0),k10_relat_1(esk2_0,X1))), inference(spm,[status(thm)],[c_0_59, c_0_60])).
fof(c_0_63, plain, ![X19, X20]:(~v1_relat_1(X20)|(~r1_tarski(k1_relat_1(X20),X19)|k7_relat_1(X20,X19)=X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t97_relat_1])])).
fof(c_0_64, plain, ![X13, X14]:(~v1_relat_1(X14)|k7_relat_1(X14,X13)=k7_relat_1(X14,k3_xboole_0(k1_relat_1(X14),X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t192_relat_1])])).
cnf(c_0_65, plain, (k10_relat_1(k7_relat_1(X1,X2),X3)=k3_xboole_0(X2,k10_relat_1(X1,X3))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_66, negated_conjecture, (k10_relat_1(esk2_0,k2_relat_1(esk2_0))=k1_relat_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_52]), c_0_53])])).
fof(c_0_67, plain, ![X37]:(~v1_relat_1(X37)|k7_relat_1(X37,k1_relat_1(X37))=X37), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t98_relat_1])])).
fof(c_0_68, plain, ![X38, X39]:(~v1_relat_1(X39)|k1_relat_1(k7_relat_1(X39,X38))=k3_xboole_0(k1_relat_1(X39),X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t90_relat_1])])).
fof(c_0_69, plain, ![X45, X46]:(~v1_relat_1(X46)|(~r1_tarski(X45,k1_relat_1(X46))|k1_relat_1(k7_relat_1(X46,X45))=X45)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t91_relat_1])])).
fof(c_0_70, plain, ![X40, X41]:(~v1_relat_1(X40)|v1_relat_1(k7_relat_1(X40,X41))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_relat_1])])).
cnf(c_0_71, plain, (k7_relat_1(X1,X2)=X1|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_72, plain, (k7_relat_1(X1,X2)=k7_relat_1(X1,k3_xboole_0(k1_relat_1(X1),X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_73, negated_conjecture, (k10_relat_1(k7_relat_1(esk2_0,X1),k2_relat_1(esk2_0))=k3_xboole_0(X1,k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_53])])).
cnf(c_0_74, plain, (k7_relat_1(X1,k1_relat_1(X1))=X1|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_75, plain, (k1_relat_1(k7_relat_1(X1,X2))=k3_xboole_0(k1_relat_1(X1),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_76, plain, (k1_relat_1(k7_relat_1(X1,X2))=X2|~v1_relat_1(X1)|~r1_tarski(X2,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_77, plain, (v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_78, plain, (k7_relat_1(X1,X2)=X1|~r1_tarski(k1_relat_1(X1),k3_xboole_0(k1_relat_1(X1),X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_79, negated_conjecture, (k3_xboole_0(k1_relat_1(esk2_0),k1_relat_1(esk2_0))=k1_relat_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_66]), c_0_53])])).
fof(c_0_80, plain, ![X26, X27]:(~v1_relat_1(X27)|r1_tarski(k1_relat_1(k7_relat_1(X27,X26)),X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t87_relat_1])])).
cnf(c_0_81, plain, (k1_relat_1(k7_relat_1(k7_relat_1(X1,X2),X3))=k3_xboole_0(X2,X3)|~r1_tarski(X2,k1_relat_1(X1))|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77])).
cnf(c_0_82, negated_conjecture, (k7_relat_1(esk2_0,k1_relat_1(esk2_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_38]), c_0_53])])).
cnf(c_0_83, plain, (r1_tarski(k1_relat_1(k7_relat_1(X1,X2)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_84, plain, (k1_relat_1(k7_relat_1(X1,X2))=k3_xboole_0(X2,X3)|~r1_tarski(k1_relat_1(k7_relat_1(X1,X2)),X3)|~r1_tarski(X2,k1_relat_1(X1))|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_71]), c_0_77])).
cnf(c_0_85, negated_conjecture, (k3_xboole_0(k1_relat_1(esk2_0),X1)=k1_relat_1(k7_relat_1(esk2_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_38]), c_0_53])])).
cnf(c_0_86, plain, (r1_tarski(k1_relat_1(k7_relat_1(X1,X2)),k3_xboole_0(k1_relat_1(X1),X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_83, c_0_72])).
cnf(c_0_87, negated_conjecture, (k1_relat_1(k7_relat_1(esk2_0,X1))=k1_relat_1(esk2_0)|~r1_tarski(k1_relat_1(esk2_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_82]), c_0_85]), c_0_38]), c_0_53])])).
cnf(c_0_88, negated_conjecture, (r1_tarski(k1_relat_1(esk2_0),k1_relat_1(k7_relat_1(esk2_0,X1)))|~r1_tarski(k1_relat_1(esk2_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_85]), c_0_53])])).
cnf(c_0_89, plain, (k3_xboole_0(X1,X2)=X2|~r1_tarski(X2,k1_relat_1(k7_relat_1(X3,X1)))|~r1_tarski(X1,k1_relat_1(X3))|~v1_relat_1(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_81]), c_0_77])).
cnf(c_0_90, negated_conjecture, (r1_tarski(esk1_0,k1_relat_1(k7_relat_1(esk2_0,X1)))|~r1_tarski(k1_relat_1(esk2_0),X1)), inference(spm,[status(thm)],[c_0_51, c_0_88])).
cnf(c_0_91, plain, (k7_relat_1(k7_relat_1(X1,X2),k3_xboole_0(X2,X3))=k7_relat_1(k7_relat_1(X1,X2),X3)|~r1_tarski(X2,k1_relat_1(X1))|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_76]), c_0_77])).
cnf(c_0_92, negated_conjecture, (k3_xboole_0(X1,esk1_0)=esk1_0|~r1_tarski(X1,k1_relat_1(esk2_0))|~r1_tarski(k1_relat_1(esk2_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_53])])).
cnf(c_0_93, negated_conjecture, (k7_relat_1(esk2_0,k1_relat_1(k7_relat_1(esk2_0,X1)))=k7_relat_1(esk2_0,X1)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_82]), c_0_38]), c_0_53])]), c_0_85])).
fof(c_0_94, plain, ![X47, X48]:(~v1_relat_1(X48)|k2_relat_1(k7_relat_1(X48,X47))=k9_relat_1(X48,X47)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t148_relat_1])])).
cnf(c_0_95, negated_conjecture, (k1_relat_1(k7_relat_1(esk2_0,esk1_0))=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_92]), c_0_38])])).
cnf(c_0_96, negated_conjecture, (v1_relat_1(k7_relat_1(esk2_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_93]), c_0_53])])).
cnf(c_0_97, plain, (k10_relat_1(k7_relat_1(X1,k1_relat_1(X2)),X3)=k1_relat_1(k7_relat_1(X2,k10_relat_1(X1,X3)))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_75, c_0_65])).
cnf(c_0_98, plain, (k7_relat_1(X1,k1_relat_1(k7_relat_1(X1,X2)))=k7_relat_1(X1,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_72, c_0_75])).
cnf(c_0_99, plain, (k2_relat_1(k7_relat_1(X1,X2))=k9_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_100, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)|~r1_tarski(X1,k1_relat_1(X3))|~v1_relat_1(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_81]), c_0_77])).
cnf(c_0_101, negated_conjecture, (r1_tarski(X1,k1_relat_1(esk2_0))|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_27, c_0_35])).
cnf(c_0_102, negated_conjecture, (k1_relat_1(k7_relat_1(k7_relat_1(esk2_0,esk1_0),X1))=k3_xboole_0(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_95]), c_0_96])])).
cnf(c_0_103, plain, (k1_relat_1(k7_relat_1(k7_relat_1(X1,X2),k10_relat_1(X1,X3)))=k10_relat_1(k7_relat_1(X1,X2),X3)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_77])).
cnf(c_0_104, plain, (k9_relat_1(X1,k3_xboole_0(k1_relat_1(X1),X2))=k2_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_99, c_0_72])).
cnf(c_0_105, negated_conjecture, (r1_tarski(k3_xboole_0(X1,X2),X2)|~r1_tarski(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_53])])).
cnf(c_0_106, negated_conjecture, (k3_xboole_0(esk1_0,k10_relat_1(esk2_0,X1))=k10_relat_1(k7_relat_1(esk2_0,esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_53])])).
cnf(c_0_107, negated_conjecture, (k9_relat_1(esk2_0,k1_relat_1(k7_relat_1(esk2_0,X1)))=k2_relat_1(k7_relat_1(esk2_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_85]), c_0_53])])).
cnf(c_0_108, negated_conjecture, (r1_tarski(k10_relat_1(k7_relat_1(esk2_0,esk1_0),X1),k10_relat_1(esk2_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_38])])).
cnf(c_0_109, negated_conjecture, (k2_relat_1(k7_relat_1(esk2_0,esk1_0))=k9_relat_1(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_107, c_0_95])).
cnf(c_0_110, negated_conjecture, (~r1_tarski(esk1_0,k10_relat_1(esk2_0,k9_relat_1(esk2_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_111, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_48]), c_0_95]), c_0_109]), c_0_96])]), c_0_110]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 112
# Proof object clause steps            : 73
# Proof object formula steps           : 39
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 50
# Proof object simplifying inferences  : 60
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 23
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 11722
# ...of these trivial                  : 53
# ...subsumed                          : 10022
# ...remaining for further processing  : 1647
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 109
# Backward-rewritten                   : 111
# Generated clauses                    : 86550
# ...of the previous two non-trivial   : 80715
# Contextual simplify-reflections      : 141
# Paramodulations                      : 86548
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1397
#    Positive orientable unit clauses  : 115
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 54
#    Non-unit-clauses                  : 1228
# Current number of unprocessed clauses: 68418
# ...number of literals in the above   : 271860
# Current number of archived formulas  : 0
# Current number of archived clauses   : 248
# Clause-clause subsumption calls (NU) : 264803
# Rec. Clause-clause subsumption calls : 116565
# Non-unit clause-clause subsumptions  : 5372
# Unit Clause-clause subsumption calls : 4838
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 488
# BW rewrite match successes           : 27
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1317561

# -------------------------------------------------
# User time                : 0.852 s
# System time              : 0.013 s
# Total time               : 0.866 s
# Maximum resident set size: 54032 pages
