# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t77_funct_2, conjecture, ![X1, X2]:((v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>(v2_funct_1(X2)<=>![X3, X4]:((r2_hidden(X3,X1)&(r2_hidden(X4,X1)&k1_funct_1(X2,X3)=k1_funct_1(X2,X4)))=>X3=X4))), file('data/MPTP/proofs/384ca8acded2451b90ad6a53d75b8531.E_input', t77_funct_2)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/384ca8acded2451b90ad6a53d75b8531.E_input', redefinition_k1_relset_1)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(((X2=k1_xboole_0=>X1=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X2,X3)))&(X2=k1_xboole_0=>(X1=k1_xboole_0|(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0))))), file('data/MPTP/proofs/384ca8acded2451b90ad6a53d75b8531.E_input', d1_funct_2)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/384ca8acded2451b90ad6a53d75b8531.E_input', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/384ca8acded2451b90ad6a53d75b8531.E_input', fc6_relat_1)).
fof(d8_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)<=>![X2, X3]:((r2_hidden(X2,k1_relat_1(X1))&(r2_hidden(X3,k1_relat_1(X1))&k1_funct_1(X1,X2)=k1_funct_1(X1,X3)))=>X2=X3))), file('data/MPTP/proofs/384ca8acded2451b90ad6a53d75b8531.E_input', d8_funct_1)).
fof(t54_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(X2=k2_funct_1(X1)<=>(k1_relat_1(X2)=k2_relat_1(X1)&![X3, X4]:(((r2_hidden(X3,k2_relat_1(X1))&X4=k1_funct_1(X2,X3))=>(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4)))&((r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4))=>(r2_hidden(X3,k2_relat_1(X1))&X4=k1_funct_1(X2,X3))))))))), file('data/MPTP/proofs/384ca8acded2451b90ad6a53d75b8531.E_input', t54_funct_1)).
fof(dt_k2_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_relat_1(k2_funct_1(X1))&v1_funct_1(k2_funct_1(X1)))), file('data/MPTP/proofs/384ca8acded2451b90ad6a53d75b8531.E_input', dt_k2_funct_1)).
fof(c_0_8, negated_conjecture, ~(![X1, X2]:((v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>(v2_funct_1(X2)<=>![X3, X4]:((r2_hidden(X3,X1)&(r2_hidden(X4,X1)&k1_funct_1(X2,X3)=k1_funct_1(X2,X4)))=>X3=X4)))), inference(assume_negation,[status(cth)],[t77_funct_2])).
fof(c_0_9, plain, ![X11, X12, X13]:(~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12)))|k1_relset_1(X11,X12,X13)=k1_relat_1(X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
fof(c_0_10, negated_conjecture, ![X9, X10]:((v1_funct_1(esk2_0)&(v1_funct_2(esk2_0,esk1_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))))&((((r2_hidden(esk3_0,esk1_0)|~v2_funct_1(esk2_0))&((r2_hidden(esk4_0,esk1_0)|~v2_funct_1(esk2_0))&(k1_funct_1(esk2_0,esk3_0)=k1_funct_1(esk2_0,esk4_0)|~v2_funct_1(esk2_0))))&(esk3_0!=esk4_0|~v2_funct_1(esk2_0)))&(v2_funct_1(esk2_0)|(~r2_hidden(X9,esk1_0)|(~r2_hidden(X10,esk1_0)|k1_funct_1(esk2_0,X9)!=k1_funct_1(esk2_0,X10))|X9=X10)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])).
fof(c_0_11, plain, ![X14, X15, X16]:((((~v1_funct_2(X16,X14,X15)|X14=k1_relset_1(X14,X15,X16)|X15=k1_xboole_0|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15))))&(X14!=k1_relset_1(X14,X15,X16)|v1_funct_2(X16,X14,X15)|X15=k1_xboole_0|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15)))))&((~v1_funct_2(X16,X14,X15)|X14=k1_relset_1(X14,X15,X16)|X14!=k1_xboole_0|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15))))&(X14!=k1_relset_1(X14,X15,X16)|v1_funct_2(X16,X14,X15)|X14!=k1_xboole_0|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15))))))&((~v1_funct_2(X16,X14,X15)|X16=k1_xboole_0|X14=k1_xboole_0|X15!=k1_xboole_0|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15))))&(X16!=k1_xboole_0|v1_funct_2(X16,X14,X15)|X14=k1_xboole_0|X15!=k1_xboole_0|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
cnf(c_0_12, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_14, plain, ![X17, X18]:(~v1_relat_1(X17)|(~m1_subset_1(X18,k1_zfmisc_1(X17))|v1_relat_1(X18))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_15, plain, ![X19, X20]:v1_relat_1(k2_zfmisc_1(X19,X20)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_16, plain, ![X21, X22, X23]:((~v2_funct_1(X21)|(~r2_hidden(X22,k1_relat_1(X21))|(~r2_hidden(X23,k1_relat_1(X21))|k1_funct_1(X21,X22)!=k1_funct_1(X21,X23))|X22=X23)|(~v1_relat_1(X21)|~v1_funct_1(X21)))&(((r2_hidden(esk5_1(X21),k1_relat_1(X21))|v2_funct_1(X21)|(~v1_relat_1(X21)|~v1_funct_1(X21)))&((r2_hidden(esk6_1(X21),k1_relat_1(X21))|v2_funct_1(X21)|(~v1_relat_1(X21)|~v1_funct_1(X21)))&(k1_funct_1(X21,esk5_1(X21))=k1_funct_1(X21,esk6_1(X21))|v2_funct_1(X21)|(~v1_relat_1(X21)|~v1_funct_1(X21)))))&(esk5_1(X21)!=esk6_1(X21)|v2_funct_1(X21)|(~v1_relat_1(X21)|~v1_funct_1(X21))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_funct_1])])])])])).
cnf(c_0_17, plain, (X2=k1_relset_1(X2,X3,X1)|X3=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, negated_conjecture, (v1_funct_2(esk2_0,esk1_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19, negated_conjecture, (k1_relset_1(esk1_0,esk1_0,esk2_0)=k1_relat_1(esk2_0)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_20, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (X2=X3|~v2_funct_1(X1)|~r2_hidden(X2,k1_relat_1(X1))|~r2_hidden(X3,k1_relat_1(X1))|k1_funct_1(X1,X2)!=k1_funct_1(X1,X3)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, negated_conjecture, (k1_relat_1(esk2_0)=esk1_0|k1_xboole_0=esk1_0), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_13]), c_0_18])]), c_0_19])).
cnf(c_0_24, negated_conjecture, (v1_relat_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_13]), c_0_21])])).
cnf(c_0_25, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_26, negated_conjecture, (v2_funct_1(esk2_0)|X1=X2|~r2_hidden(X1,esk1_0)|~r2_hidden(X2,esk1_0)|k1_funct_1(esk2_0,X1)!=k1_funct_1(esk2_0,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_27, plain, (k1_funct_1(X1,esk5_1(X1))=k1_funct_1(X1,esk6_1(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28, negated_conjecture, (k1_xboole_0=esk1_0|X1=X2|k1_funct_1(esk2_0,X1)!=k1_funct_1(esk2_0,X2)|~r2_hidden(X2,esk1_0)|~r2_hidden(X1,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_29, negated_conjecture, (k1_funct_1(esk2_0,esk3_0)=k1_funct_1(esk2_0,esk4_0)|~v2_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_30, negated_conjecture, (r2_hidden(esk4_0,esk1_0)|~v2_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_31, negated_conjecture, (esk6_1(esk2_0)=X1|v2_funct_1(esk2_0)|k1_funct_1(esk2_0,esk5_1(esk2_0))!=k1_funct_1(esk2_0,X1)|~r2_hidden(esk6_1(esk2_0),esk1_0)|~r2_hidden(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_25])]), c_0_24])])).
cnf(c_0_32, plain, (r2_hidden(esk6_1(X1),k1_relat_1(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_33, plain, (r2_hidden(esk5_1(X1),k1_relat_1(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_34, negated_conjecture, (k1_xboole_0=esk1_0|X1=esk4_0|k1_funct_1(esk2_0,X1)!=k1_funct_1(esk2_0,esk3_0)|~r2_hidden(X1,esk1_0)|~v2_funct_1(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])).
cnf(c_0_35, negated_conjecture, (r2_hidden(esk3_0,esk1_0)|~v2_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_36, negated_conjecture, (esk3_0!=esk4_0|~v2_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_37, plain, ![X26, X27, X28, X29, X30, X31]:(((k1_relat_1(X27)=k2_relat_1(X26)|X27!=k2_funct_1(X26)|(~v1_relat_1(X27)|~v1_funct_1(X27))|~v2_funct_1(X26)|(~v1_relat_1(X26)|~v1_funct_1(X26)))&(((r2_hidden(X29,k1_relat_1(X26))|(~r2_hidden(X28,k2_relat_1(X26))|X29!=k1_funct_1(X27,X28))|X27!=k2_funct_1(X26)|(~v1_relat_1(X27)|~v1_funct_1(X27))|~v2_funct_1(X26)|(~v1_relat_1(X26)|~v1_funct_1(X26)))&(X28=k1_funct_1(X26,X29)|(~r2_hidden(X28,k2_relat_1(X26))|X29!=k1_funct_1(X27,X28))|X27!=k2_funct_1(X26)|(~v1_relat_1(X27)|~v1_funct_1(X27))|~v2_funct_1(X26)|(~v1_relat_1(X26)|~v1_funct_1(X26))))&((r2_hidden(X30,k2_relat_1(X26))|(~r2_hidden(X31,k1_relat_1(X26))|X30!=k1_funct_1(X26,X31))|X27!=k2_funct_1(X26)|(~v1_relat_1(X27)|~v1_funct_1(X27))|~v2_funct_1(X26)|(~v1_relat_1(X26)|~v1_funct_1(X26)))&(X31=k1_funct_1(X27,X30)|(~r2_hidden(X31,k1_relat_1(X26))|X30!=k1_funct_1(X26,X31))|X27!=k2_funct_1(X26)|(~v1_relat_1(X27)|~v1_funct_1(X27))|~v2_funct_1(X26)|(~v1_relat_1(X26)|~v1_funct_1(X26))))))&(((((r2_hidden(esk10_2(X26,X27),k1_relat_1(X26))|r2_hidden(esk7_2(X26,X27),k2_relat_1(X26))|k1_relat_1(X27)!=k2_relat_1(X26)|X27=k2_funct_1(X26)|(~v1_relat_1(X27)|~v1_funct_1(X27))|~v2_funct_1(X26)|(~v1_relat_1(X26)|~v1_funct_1(X26)))&(esk9_2(X26,X27)=k1_funct_1(X26,esk10_2(X26,X27))|r2_hidden(esk7_2(X26,X27),k2_relat_1(X26))|k1_relat_1(X27)!=k2_relat_1(X26)|X27=k2_funct_1(X26)|(~v1_relat_1(X27)|~v1_funct_1(X27))|~v2_funct_1(X26)|(~v1_relat_1(X26)|~v1_funct_1(X26))))&(~r2_hidden(esk9_2(X26,X27),k2_relat_1(X26))|esk10_2(X26,X27)!=k1_funct_1(X27,esk9_2(X26,X27))|r2_hidden(esk7_2(X26,X27),k2_relat_1(X26))|k1_relat_1(X27)!=k2_relat_1(X26)|X27=k2_funct_1(X26)|(~v1_relat_1(X27)|~v1_funct_1(X27))|~v2_funct_1(X26)|(~v1_relat_1(X26)|~v1_funct_1(X26))))&(((r2_hidden(esk10_2(X26,X27),k1_relat_1(X26))|esk8_2(X26,X27)=k1_funct_1(X27,esk7_2(X26,X27))|k1_relat_1(X27)!=k2_relat_1(X26)|X27=k2_funct_1(X26)|(~v1_relat_1(X27)|~v1_funct_1(X27))|~v2_funct_1(X26)|(~v1_relat_1(X26)|~v1_funct_1(X26)))&(esk9_2(X26,X27)=k1_funct_1(X26,esk10_2(X26,X27))|esk8_2(X26,X27)=k1_funct_1(X27,esk7_2(X26,X27))|k1_relat_1(X27)!=k2_relat_1(X26)|X27=k2_funct_1(X26)|(~v1_relat_1(X27)|~v1_funct_1(X27))|~v2_funct_1(X26)|(~v1_relat_1(X26)|~v1_funct_1(X26))))&(~r2_hidden(esk9_2(X26,X27),k2_relat_1(X26))|esk10_2(X26,X27)!=k1_funct_1(X27,esk9_2(X26,X27))|esk8_2(X26,X27)=k1_funct_1(X27,esk7_2(X26,X27))|k1_relat_1(X27)!=k2_relat_1(X26)|X27=k2_funct_1(X26)|(~v1_relat_1(X27)|~v1_funct_1(X27))|~v2_funct_1(X26)|(~v1_relat_1(X26)|~v1_funct_1(X26)))))&(((r2_hidden(esk10_2(X26,X27),k1_relat_1(X26))|(~r2_hidden(esk8_2(X26,X27),k1_relat_1(X26))|esk7_2(X26,X27)!=k1_funct_1(X26,esk8_2(X26,X27)))|k1_relat_1(X27)!=k2_relat_1(X26)|X27=k2_funct_1(X26)|(~v1_relat_1(X27)|~v1_funct_1(X27))|~v2_funct_1(X26)|(~v1_relat_1(X26)|~v1_funct_1(X26)))&(esk9_2(X26,X27)=k1_funct_1(X26,esk10_2(X26,X27))|(~r2_hidden(esk8_2(X26,X27),k1_relat_1(X26))|esk7_2(X26,X27)!=k1_funct_1(X26,esk8_2(X26,X27)))|k1_relat_1(X27)!=k2_relat_1(X26)|X27=k2_funct_1(X26)|(~v1_relat_1(X27)|~v1_funct_1(X27))|~v2_funct_1(X26)|(~v1_relat_1(X26)|~v1_funct_1(X26))))&(~r2_hidden(esk9_2(X26,X27),k2_relat_1(X26))|esk10_2(X26,X27)!=k1_funct_1(X27,esk9_2(X26,X27))|(~r2_hidden(esk8_2(X26,X27),k1_relat_1(X26))|esk7_2(X26,X27)!=k1_funct_1(X26,esk8_2(X26,X27)))|k1_relat_1(X27)!=k2_relat_1(X26)|X27=k2_funct_1(X26)|(~v1_relat_1(X27)|~v1_funct_1(X27))|~v2_funct_1(X26)|(~v1_relat_1(X26)|~v1_funct_1(X26)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t54_funct_1])])])])])])).
fof(c_0_38, plain, ![X36]:((v1_relat_1(k2_funct_1(X36))|(~v1_relat_1(X36)|~v1_funct_1(X36)))&(v1_funct_1(k2_funct_1(X36))|(~v1_relat_1(X36)|~v1_funct_1(X36)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_funct_1])])])).
cnf(c_0_39, negated_conjecture, (esk6_1(esk2_0)=esk5_1(esk2_0)|v2_funct_1(esk2_0)|~r2_hidden(esk6_1(esk2_0),esk1_0)|~r2_hidden(esk5_1(esk2_0),esk1_0)), inference(er,[status(thm)],[c_0_31])).
cnf(c_0_40, negated_conjecture, (k1_xboole_0=esk1_0|r2_hidden(esk6_1(esk2_0),esk1_0)|v2_funct_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_23]), c_0_24]), c_0_25])])).
cnf(c_0_41, negated_conjecture, (k1_xboole_0=esk1_0|r2_hidden(esk5_1(esk2_0),esk1_0)|v2_funct_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_23]), c_0_24]), c_0_25])])).
cnf(c_0_42, negated_conjecture, (k1_xboole_0=esk1_0|~v2_funct_1(esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_34]), c_0_35]), c_0_36])).
cnf(c_0_43, plain, (X1=k1_funct_1(X2,X3)|~r2_hidden(X1,k1_relat_1(X4))|X3!=k1_funct_1(X4,X1)|X2!=k2_funct_1(X4)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X4)|~v1_relat_1(X4)|~v1_funct_1(X4)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_44, plain, (v1_funct_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_45, plain, (v1_relat_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_46, plain, (X2=k1_relset_1(X2,X3,X1)|~v1_funct_2(X1,X2,X3)|X2!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_47, plain, (v2_funct_1(X1)|esk5_1(X1)!=esk6_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_48, negated_conjecture, (esk6_1(esk2_0)=esk5_1(esk2_0)|k1_xboole_0=esk1_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41]), c_0_42])).
cnf(c_0_49, plain, (k1_funct_1(k2_funct_1(X1),k1_funct_1(X1,X2))=X2|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_43])]), c_0_44]), c_0_45])).
cnf(c_0_50, plain, (k1_relset_1(k1_xboole_0,X1,X2)=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,X1)))|~v1_funct_2(X2,k1_xboole_0,X1)), inference(er,[status(thm)],[c_0_46])).
cnf(c_0_51, negated_conjecture, (k1_xboole_0=esk1_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_24]), c_0_25])]), c_0_42])).
cnf(c_0_52, negated_conjecture, (k1_funct_1(k2_funct_1(esk2_0),k1_funct_1(esk2_0,esk3_0))=esk4_0|~r2_hidden(esk4_0,k1_relat_1(esk2_0))|~v2_funct_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_29]), c_0_24]), c_0_25])])).
cnf(c_0_53, plain, (k1_relset_1(esk1_0,X1,X2)=esk1_0|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,X1)))|~v1_funct_2(X2,esk1_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_51]), c_0_51]), c_0_51]), c_0_51])).
cnf(c_0_54, negated_conjecture, (~r2_hidden(esk3_0,k1_relat_1(esk2_0))|~r2_hidden(esk4_0,k1_relat_1(esk2_0))|~v2_funct_1(esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_52]), c_0_24]), c_0_25])]), c_0_36])).
cnf(c_0_55, negated_conjecture, (k1_relat_1(esk2_0)=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_13]), c_0_19]), c_0_18])])).
cnf(c_0_56, negated_conjecture, (~v2_funct_1(esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_55]), c_0_55]), c_0_30]), c_0_35])).
cnf(c_0_57, negated_conjecture, (r2_hidden(esk5_1(esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_55]), c_0_24]), c_0_25])]), c_0_56])).
cnf(c_0_58, negated_conjecture, (esk6_1(esk2_0)=esk5_1(esk2_0)|~r2_hidden(esk6_1(esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_57])]), c_0_56])).
cnf(c_0_59, negated_conjecture, (r2_hidden(esk6_1(esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_55]), c_0_24]), c_0_25])]), c_0_56])).
cnf(c_0_60, negated_conjecture, (esk6_1(esk2_0)=esk5_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_59])])).
cnf(c_0_61, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_60]), c_0_24]), c_0_25])]), c_0_56]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 62
# Proof object clause steps            : 45
# Proof object formula steps           : 17
# Proof object conjectures             : 32
# Proof object clause conjectures      : 29
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 68
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 38
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 38
# Processed clauses                    : 76
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 74
# Other redundant clauses eliminated   : 17
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 20
# Generated clauses                    : 106
# ...of the previous two non-trivial   : 94
# Contextual simplify-reflections      : 20
# Paramodulations                      : 90
# Factorizations                       : 0
# Equation resolutions                 : 21
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 44
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 33
# Current number of unprocessed clauses: 51
# ...number of literals in the above   : 427
# Current number of archived formulas  : 0
# Current number of archived clauses   : 21
# Clause-clause subsumption calls (NU) : 756
# Rec. Clause-clause subsumption calls : 83
# Non-unit clause-clause subsumptions  : 23
# Unit Clause-clause subsumption calls : 10
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5922

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28628 pages
