# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t5_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1)&v2_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1)))), file('data/MPTP/proofs/3857a728719b4ed486a421c44a0e506b.E_input', t5_waybel_0)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/3857a728719b4ed486a421c44a0e506b.E_input', d2_tarski)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/3857a728719b4ed486a421c44a0e506b.E_input', redefinition_k6_domain_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/3857a728719b4ed486a421c44a0e506b.E_input', t69_enumset1)).
fof(d2_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v2_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X3,X2)&(r2_hidden(X4,X2)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>~((r2_hidden(X5,X2)&(r1_orders_2(X1,X5,X3)&r1_orders_2(X1,X5,X4))))))))))))), file('data/MPTP/proofs/3857a728719b4ed486a421c44a0e506b.E_input', d2_waybel_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/3857a728719b4ed486a421c44a0e506b.E_input', t4_subset)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/3857a728719b4ed486a421c44a0e506b.E_input', t24_orders_2)).
fof(t63_subset_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))), file('data/MPTP/proofs/3857a728719b4ed486a421c44a0e506b.E_input', t63_subset_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/3857a728719b4ed486a421c44a0e506b.E_input', t2_subset)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/3857a728719b4ed486a421c44a0e506b.E_input', fc2_struct_0)).
fof(d1_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v1_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X3,X2)&(r2_hidden(X4,X2)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>~((r2_hidden(X5,X2)&(r1_orders_2(X1,X3,X5)&r1_orders_2(X1,X4,X5))))))))))))), file('data/MPTP/proofs/3857a728719b4ed486a421c44a0e506b.E_input', d1_waybel_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/3857a728719b4ed486a421c44a0e506b.E_input', dt_l1_orders_2)).
fof(c_0_12, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1)&v2_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1))))), inference(assume_negation,[status(cth)],[t5_waybel_0])).
fof(c_0_13, plain, ![X8, X9, X10, X11, X12, X13, X14, X15]:(((~r2_hidden(X11,X10)|(X11=X8|X11=X9)|X10!=k2_tarski(X8,X9))&((X12!=X8|r2_hidden(X12,X10)|X10!=k2_tarski(X8,X9))&(X12!=X9|r2_hidden(X12,X10)|X10!=k2_tarski(X8,X9))))&(((esk3_3(X13,X14,X15)!=X13|~r2_hidden(esk3_3(X13,X14,X15),X15)|X15=k2_tarski(X13,X14))&(esk3_3(X13,X14,X15)!=X14|~r2_hidden(esk3_3(X13,X14,X15),X15)|X15=k2_tarski(X13,X14)))&(r2_hidden(esk3_3(X13,X14,X15),X15)|(esk3_3(X13,X14,X15)=X13|esk3_3(X13,X14,X15)=X14)|X15=k2_tarski(X13,X14)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_14, plain, ![X17, X18]:(v1_xboole_0(X17)|~m1_subset_1(X18,X17)|k6_domain_1(X17,X18)=k1_tarski(X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_15, plain, ![X19]:k2_tarski(X19,X19)=k1_tarski(X19), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_16, plain, ![X35, X36, X37, X38, X42]:(((m1_subset_1(esk7_4(X35,X36,X37,X38),u1_struct_0(X35))|~r2_hidden(X38,X36)|~r2_hidden(X37,X36)|~m1_subset_1(X38,u1_struct_0(X35))|~m1_subset_1(X37,u1_struct_0(X35))|~v2_waybel_0(X36,X35)|~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X35)))|~l1_orders_2(X35))&((r2_hidden(esk7_4(X35,X36,X37,X38),X36)|~r2_hidden(X38,X36)|~r2_hidden(X37,X36)|~m1_subset_1(X38,u1_struct_0(X35))|~m1_subset_1(X37,u1_struct_0(X35))|~v2_waybel_0(X36,X35)|~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X35)))|~l1_orders_2(X35))&((r1_orders_2(X35,esk7_4(X35,X36,X37,X38),X37)|~r2_hidden(X38,X36)|~r2_hidden(X37,X36)|~m1_subset_1(X38,u1_struct_0(X35))|~m1_subset_1(X37,u1_struct_0(X35))|~v2_waybel_0(X36,X35)|~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X35)))|~l1_orders_2(X35))&(r1_orders_2(X35,esk7_4(X35,X36,X37,X38),X38)|~r2_hidden(X38,X36)|~r2_hidden(X37,X36)|~m1_subset_1(X38,u1_struct_0(X35))|~m1_subset_1(X37,u1_struct_0(X35))|~v2_waybel_0(X36,X35)|~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X35)))|~l1_orders_2(X35)))))&((m1_subset_1(esk8_2(X35,X36),u1_struct_0(X35))|v2_waybel_0(X36,X35)|~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X35)))|~l1_orders_2(X35))&((m1_subset_1(esk9_2(X35,X36),u1_struct_0(X35))|v2_waybel_0(X36,X35)|~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X35)))|~l1_orders_2(X35))&((r2_hidden(esk8_2(X35,X36),X36)|v2_waybel_0(X36,X35)|~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X35)))|~l1_orders_2(X35))&((r2_hidden(esk9_2(X35,X36),X36)|v2_waybel_0(X36,X35)|~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X35)))|~l1_orders_2(X35))&(~m1_subset_1(X42,u1_struct_0(X35))|(~r2_hidden(X42,X36)|(~r1_orders_2(X35,X42,esk8_2(X35,X36))|~r1_orders_2(X35,X42,esk9_2(X35,X36))))|v2_waybel_0(X36,X35)|~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X35)))|~l1_orders_2(X35))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_waybel_0])])])])])).
fof(c_0_17, plain, ![X28, X29, X30]:(~r2_hidden(X28,X29)|~m1_subset_1(X29,k1_zfmisc_1(X30))|m1_subset_1(X28,X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_18, plain, ![X31, X32]:(v2_struct_0(X31)|(~v3_orders_2(X31)|~l1_orders_2(X31))|(~m1_subset_1(X32,u1_struct_0(X31))|r1_orders_2(X31,X32,X32))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_19, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&l1_orders_2(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(~v1_waybel_0(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)|~v2_waybel_0(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])).
cnf(c_0_20, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (v2_waybel_0(X3,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r1_orders_2(X2,X1,esk8_2(X2,X3))|~r1_orders_2(X2,X1,esk9_2(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (X1=X2|X1=X3|~r2_hidden(X1,k2_tarski(X3,X2))), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_30, plain, (k2_tarski(X2,X2)=k6_domain_1(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_31, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_32, plain, (v2_waybel_0(X1,X2)|~r1_orders_2(X2,X3,esk8_2(X2,X1))|~r1_orders_2(X2,X3,esk9_2(X2,X1))|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(csr,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_33, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])]), c_0_28])).
cnf(c_0_34, plain, (X1=X2|v1_xboole_0(X3)|~r2_hidden(X1,k6_domain_1(X3,X2))|~m1_subset_1(X2,X3)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_35, plain, (r2_hidden(esk9_2(X1,X2),X2)|v2_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_36, plain, (r2_hidden(X1,k2_tarski(X1,X2))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_31])])).
fof(c_0_37, plain, ![X33, X34]:(~r2_hidden(X33,X34)|m1_subset_1(k1_tarski(X33),k1_zfmisc_1(X34))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_subset_1])])).
cnf(c_0_38, negated_conjecture, (v2_waybel_0(X1,esk1_0)|~r1_orders_2(esk1_0,esk9_2(esk1_0,X1),esk8_2(esk1_0,X1))|~r2_hidden(esk9_2(esk1_0,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_27])]), c_0_24])).
cnf(c_0_39, plain, (esk9_2(X1,k6_domain_1(X2,X3))=X3|v1_xboole_0(X2)|v2_waybel_0(k6_domain_1(X2,X3),X1)|~m1_subset_1(k6_domain_1(X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40, plain, (v1_xboole_0(X1)|r2_hidden(X2,k6_domain_1(X1,X2))|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_36, c_0_30])).
cnf(c_0_41, plain, (r2_hidden(esk8_2(X1,X2),X2)|v2_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_42, plain, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_43, plain, ![X43, X44]:(~m1_subset_1(X43,X44)|(v1_xboole_0(X44)|r2_hidden(X43,X44))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_44, negated_conjecture, (v1_xboole_0(X1)|v2_waybel_0(k6_domain_1(X1,X2),esk1_0)|~r1_orders_2(esk1_0,X2,esk8_2(esk1_0,k6_domain_1(X1,X2)))|~m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_27])]), c_0_40])).
cnf(c_0_45, plain, (esk8_2(X1,k6_domain_1(X2,X3))=X3|v1_xboole_0(X2)|v2_waybel_0(k6_domain_1(X2,X3),X1)|~m1_subset_1(k6_domain_1(X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_34, c_0_41])).
cnf(c_0_46, plain, (m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(X2))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_42, c_0_22])).
cnf(c_0_47, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_48, plain, ![X45]:(v2_struct_0(X45)|~l1_struct_0(X45)|~v1_xboole_0(u1_struct_0(X45))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_49, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_50, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_51, plain, ![X20, X21, X22, X23, X27]:(((m1_subset_1(esk4_4(X20,X21,X22,X23),u1_struct_0(X20))|~r2_hidden(X23,X21)|~r2_hidden(X22,X21)|~m1_subset_1(X23,u1_struct_0(X20))|~m1_subset_1(X22,u1_struct_0(X20))|~v1_waybel_0(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_orders_2(X20))&((r2_hidden(esk4_4(X20,X21,X22,X23),X21)|~r2_hidden(X23,X21)|~r2_hidden(X22,X21)|~m1_subset_1(X23,u1_struct_0(X20))|~m1_subset_1(X22,u1_struct_0(X20))|~v1_waybel_0(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_orders_2(X20))&((r1_orders_2(X20,X22,esk4_4(X20,X21,X22,X23))|~r2_hidden(X23,X21)|~r2_hidden(X22,X21)|~m1_subset_1(X23,u1_struct_0(X20))|~m1_subset_1(X22,u1_struct_0(X20))|~v1_waybel_0(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_orders_2(X20))&(r1_orders_2(X20,X23,esk4_4(X20,X21,X22,X23))|~r2_hidden(X23,X21)|~r2_hidden(X22,X21)|~m1_subset_1(X23,u1_struct_0(X20))|~m1_subset_1(X22,u1_struct_0(X20))|~v1_waybel_0(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_orders_2(X20)))))&((m1_subset_1(esk5_2(X20,X21),u1_struct_0(X20))|v1_waybel_0(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_orders_2(X20))&((m1_subset_1(esk6_2(X20,X21),u1_struct_0(X20))|v1_waybel_0(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_orders_2(X20))&((r2_hidden(esk5_2(X20,X21),X21)|v1_waybel_0(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_orders_2(X20))&((r2_hidden(esk6_2(X20,X21),X21)|v1_waybel_0(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_orders_2(X20))&(~m1_subset_1(X27,u1_struct_0(X20))|(~r2_hidden(X27,X21)|(~r1_orders_2(X20,esk5_2(X20,X21),X27)|~r1_orders_2(X20,esk6_2(X20,X21),X27)))|v1_waybel_0(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_orders_2(X20))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_waybel_0])])])])])).
cnf(c_0_52, negated_conjecture, (v1_xboole_0(X1)|v2_waybel_0(k6_domain_1(X1,X2),esk1_0)|~r1_orders_2(esk1_0,X2,X2)|~m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_27])])).
cnf(c_0_53, plain, (v1_xboole_0(X1)|m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X3))|~r2_hidden(X2,X3)|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_46, c_0_30])).
cnf(c_0_54, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k2_tarski(X3,X3))|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_24, c_0_46])).
cnf(c_0_55, plain, (r2_hidden(X1,k2_tarski(X2,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_47])])).
cnf(c_0_56, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_57, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
fof(c_0_58, plain, ![X46]:(~l1_orders_2(X46)|l1_struct_0(X46)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_59, plain, (v1_waybel_0(X3,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r1_orders_2(X2,esk5_2(X2,X3),X1)|~r1_orders_2(X2,esk6_2(X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_60, plain, (r2_hidden(esk6_2(X1,X2),X2)|v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_61, negated_conjecture, (v1_xboole_0(X1)|v2_waybel_0(k6_domain_1(X1,X2),esk1_0)|~r1_orders_2(esk1_0,X2,X2)|~r2_hidden(X2,u1_struct_0(esk1_0))|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_62, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_63, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_28])).
cnf(c_0_64, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_65, plain, (v1_waybel_0(X1,X2)|~r1_orders_2(X2,esk5_2(X2,X1),X3)|~r1_orders_2(X2,esk6_2(X2,X1),X3)|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(csr,[status(thm)],[c_0_59, c_0_24])).
cnf(c_0_66, plain, (esk6_2(X1,k2_tarski(X2,X3))=X3|esk6_2(X1,k2_tarski(X2,X3))=X2|v1_waybel_0(k2_tarski(X2,X3),X1)|~m1_subset_1(k2_tarski(X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_29, c_0_60])).
cnf(c_0_67, plain, (r2_hidden(esk5_2(X1,X2),X2)|v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_68, negated_conjecture, (~v1_waybel_0(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)|~v2_waybel_0(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_69, negated_conjecture, (v1_xboole_0(X1)|v2_waybel_0(k6_domain_1(X1,X2),esk1_0)|~r2_hidden(X2,u1_struct_0(esk1_0))|~m1_subset_1(X2,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_33]), c_0_62])).
cnf(c_0_70, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_27])])).
cnf(c_0_71, negated_conjecture, (v1_waybel_0(X1,esk1_0)|~r1_orders_2(esk1_0,esk5_2(esk1_0,X1),esk6_2(esk1_0,X1))|~r2_hidden(esk6_2(esk1_0,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_33]), c_0_27])]), c_0_24])).
cnf(c_0_72, plain, (esk6_2(X1,k2_tarski(X2,X2))=X2|v1_waybel_0(k2_tarski(X2,X2),X1)|~r2_hidden(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_66, c_0_46])).
cnf(c_0_73, plain, (esk5_2(X1,k2_tarski(X2,X3))=X3|esk5_2(X1,k2_tarski(X2,X3))=X2|v1_waybel_0(k2_tarski(X2,X3),X1)|~m1_subset_1(k2_tarski(X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_29, c_0_67])).
cnf(c_0_74, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~v1_waybel_0(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70]), c_0_50])])).
cnf(c_0_75, negated_conjecture, (v1_waybel_0(k2_tarski(X1,X1),esk1_0)|~r1_orders_2(esk1_0,esk5_2(esk1_0,k2_tarski(X1,X1)),X1)|~r2_hidden(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_36]), c_0_27])]), c_0_46])).
cnf(c_0_76, plain, (esk5_2(X1,k2_tarski(X2,X2))=X2|v1_waybel_0(k2_tarski(X2,X2),X1)|~r2_hidden(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_73, c_0_46])).
cnf(c_0_77, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~v1_waybel_0(k2_tarski(esk2_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_30]), c_0_50])])).
cnf(c_0_78, negated_conjecture, (v1_waybel_0(k2_tarski(X1,X1),esk1_0)|~r1_orders_2(esk1_0,X1,X1)|~r2_hidden(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_27])])).
cnf(c_0_79, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_70])])).
cnf(c_0_80, negated_conjecture, (~l1_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_79]), c_0_28])).
cnf(c_0_81, negated_conjecture, (~r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_64]), c_0_27])])).
cnf(c_0_82, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_33]), c_0_50])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 83
# Proof object clause steps            : 58
# Proof object formula steps           : 25
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 29
# Proof object simplifying inferences  : 45
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 37
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 36
# Processed clauses                    : 207
# ...of these trivial                  : 0
# ...subsumed                          : 64
# ...remaining for further processing  : 143
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 2
# Generated clauses                    : 312
# ...of the previous two non-trivial   : 298
# Contextual simplify-reflections      : 23
# Paramodulations                      : 309
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 99
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 89
# Current number of unprocessed clauses: 163
# ...number of literals in the above   : 1098
# Current number of archived formulas  : 0
# Current number of archived clauses   : 42
# Clause-clause subsumption calls (NU) : 3775
# Rec. Clause-clause subsumption calls : 678
# Non-unit clause-clause subsumptions  : 90
# Unit Clause-clause subsumption calls : 22
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 10644

# -------------------------------------------------
# User time                : 0.017 s
# System time              : 0.000 s
# Total time               : 0.017 s
# Maximum resident set size: 28672 pages
