# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_funct_2, axiom, ![X1, X2, X3]:(X3=k1_funct_2(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:((v1_relat_1(X5)&v1_funct_1(X5))&(X4=X5&(k1_relat_1(X5)=X1&r1_tarski(k2_relat_1(X5),X2)))))), file('data/MPTP/proofs/3860d8c5059e4688af9dbe55b4d127a0.E_input', d2_funct_2)).
fof(t121_funct_2, conjecture, ![X1, X2, X3]:(r2_hidden(X3,k1_funct_2(X1,X2))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))), file('data/MPTP/proofs/3860d8c5059e4688af9dbe55b4d127a0.E_input', t121_funct_2)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/3860d8c5059e4688af9dbe55b4d127a0.E_input', d3_tarski)).
fof(t12_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r2_hidden(X1,k1_relat_1(X2))=>r2_hidden(k1_funct_1(X2,X1),k2_relat_1(X2)))), file('data/MPTP/proofs/3860d8c5059e4688af9dbe55b4d127a0.E_input', t12_funct_1)).
fof(t5_funct_2, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X3)=X1&![X4]:(r2_hidden(X4,X1)=>r2_hidden(k1_funct_1(X3,X4),X2)))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))), file('data/MPTP/proofs/3860d8c5059e4688af9dbe55b4d127a0.E_input', t5_funct_2)).
fof(c_0_5, plain, ![X11, X12, X13, X14, X16, X17, X18, X19, X20, X22]:(((((v1_relat_1(esk4_4(X11,X12,X13,X14))|~r2_hidden(X14,X13)|X13!=k1_funct_2(X11,X12))&(v1_funct_1(esk4_4(X11,X12,X13,X14))|~r2_hidden(X14,X13)|X13!=k1_funct_2(X11,X12)))&((X14=esk4_4(X11,X12,X13,X14)|~r2_hidden(X14,X13)|X13!=k1_funct_2(X11,X12))&((k1_relat_1(esk4_4(X11,X12,X13,X14))=X11|~r2_hidden(X14,X13)|X13!=k1_funct_2(X11,X12))&(r1_tarski(k2_relat_1(esk4_4(X11,X12,X13,X14)),X12)|~r2_hidden(X14,X13)|X13!=k1_funct_2(X11,X12)))))&(~v1_relat_1(X17)|~v1_funct_1(X17)|(X16!=X17|(k1_relat_1(X17)!=X11|~r1_tarski(k2_relat_1(X17),X12)))|r2_hidden(X16,X13)|X13!=k1_funct_2(X11,X12)))&((~r2_hidden(esk5_3(X18,X19,X20),X20)|(~v1_relat_1(X22)|~v1_funct_1(X22)|(esk5_3(X18,X19,X20)!=X22|(k1_relat_1(X22)!=X18|~r1_tarski(k2_relat_1(X22),X19))))|X20=k1_funct_2(X18,X19))&(((v1_relat_1(esk6_3(X18,X19,X20))|r2_hidden(esk5_3(X18,X19,X20),X20)|X20=k1_funct_2(X18,X19))&(v1_funct_1(esk6_3(X18,X19,X20))|r2_hidden(esk5_3(X18,X19,X20),X20)|X20=k1_funct_2(X18,X19)))&((esk5_3(X18,X19,X20)=esk6_3(X18,X19,X20)|r2_hidden(esk5_3(X18,X19,X20),X20)|X20=k1_funct_2(X18,X19))&((k1_relat_1(esk6_3(X18,X19,X20))=X18|r2_hidden(esk5_3(X18,X19,X20),X20)|X20=k1_funct_2(X18,X19))&(r1_tarski(k2_relat_1(esk6_3(X18,X19,X20)),X19)|r2_hidden(esk5_3(X18,X19,X20),X20)|X20=k1_funct_2(X18,X19))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_funct_2])])])])])])).
fof(c_0_6, negated_conjecture, ~(![X1, X2, X3]:(r2_hidden(X3,k1_funct_2(X1,X2))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))), inference(assume_negation,[status(cth)],[t121_funct_2])).
cnf(c_0_7, plain, (X1=esk4_4(X2,X3,X4,X1)|~r2_hidden(X1,X4)|X4!=k1_funct_2(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_8, negated_conjecture, (r2_hidden(esk3_0,k1_funct_2(esk1_0,esk2_0))&(~v1_funct_1(esk3_0)|(~v1_funct_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
cnf(c_0_9, plain, (r1_tarski(k2_relat_1(esk4_4(X1,X2,X3,X4)),X2)|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10, plain, (esk4_4(X1,X2,k1_funct_2(X1,X2),X3)=X3|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_7])).
cnf(c_0_11, negated_conjecture, (r2_hidden(esk3_0,k1_funct_2(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_12, plain, ![X24, X25, X26, X27, X28]:((~r1_tarski(X24,X25)|(~r2_hidden(X26,X24)|r2_hidden(X26,X25)))&((r2_hidden(esk7_2(X27,X28),X27)|r1_tarski(X27,X28))&(~r2_hidden(esk7_2(X27,X28),X28)|r1_tarski(X27,X28)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_13, plain, (r1_tarski(k2_relat_1(esk4_4(X1,X2,k1_funct_2(X1,X2),X3)),X2)|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_9])).
cnf(c_0_14, negated_conjecture, (esk4_4(esk1_0,esk2_0,k1_funct_2(esk1_0,esk2_0),esk3_0)=esk3_0), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15, plain, (v1_relat_1(esk4_4(X1,X2,X3,X4))|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_16, plain, (k1_relat_1(esk4_4(X1,X2,X3,X4))=X1|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_17, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, negated_conjecture, (r1_tarski(k2_relat_1(esk3_0),esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_11]), c_0_14])).
fof(c_0_19, plain, ![X9, X10]:(~v1_relat_1(X10)|~v1_funct_1(X10)|(~r2_hidden(X9,k1_relat_1(X10))|r2_hidden(k1_funct_1(X10,X9),k2_relat_1(X10)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_funct_1])])).
cnf(c_0_20, plain, (v1_relat_1(esk4_4(X1,X2,k1_funct_2(X1,X2),X3))|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (v1_funct_1(esk4_4(X1,X2,X3,X4))|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_22, plain, (k1_relat_1(esk4_4(X1,X2,k1_funct_2(X1,X2),X3))=X1|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_16])).
fof(c_0_23, plain, ![X30, X31, X32]:(((v1_funct_1(X32)|(r2_hidden(esk8_3(X30,X31,X32),X30)|k1_relat_1(X32)!=X30)|(~v1_relat_1(X32)|~v1_funct_1(X32)))&((v1_funct_2(X32,X30,X31)|(r2_hidden(esk8_3(X30,X31,X32),X30)|k1_relat_1(X32)!=X30)|(~v1_relat_1(X32)|~v1_funct_1(X32)))&(m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(X30,X31)))|(r2_hidden(esk8_3(X30,X31,X32),X30)|k1_relat_1(X32)!=X30)|(~v1_relat_1(X32)|~v1_funct_1(X32)))))&((v1_funct_1(X32)|(~r2_hidden(k1_funct_1(X32,esk8_3(X30,X31,X32)),X31)|k1_relat_1(X32)!=X30)|(~v1_relat_1(X32)|~v1_funct_1(X32)))&((v1_funct_2(X32,X30,X31)|(~r2_hidden(k1_funct_1(X32,esk8_3(X30,X31,X32)),X31)|k1_relat_1(X32)!=X30)|(~v1_relat_1(X32)|~v1_funct_1(X32)))&(m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(X30,X31)))|(~r2_hidden(k1_funct_1(X32,esk8_3(X30,X31,X32)),X31)|k1_relat_1(X32)!=X30)|(~v1_relat_1(X32)|~v1_funct_1(X32)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_funct_2])])])])).
cnf(c_0_24, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,k2_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_25, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, negated_conjecture, (v1_relat_1(esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_11]), c_0_14])).
cnf(c_0_27, negated_conjecture, (v1_funct_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_14]), c_0_11])])).
cnf(c_0_28, negated_conjecture, (k1_relat_1(esk3_0)=esk1_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_11]), c_0_14])).
cnf(c_0_29, plain, (v1_funct_2(X1,X2,X3)|~r2_hidden(k1_funct_1(X1,esk8_3(X2,X3,X1)),X3)|k1_relat_1(X1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,X1),esk2_0)|~r2_hidden(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26]), c_0_27]), c_0_28])])).
cnf(c_0_31, negated_conjecture, (~v1_funct_1(esk3_0)|~v1_funct_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_32, negated_conjecture, (v1_funct_2(esk3_0,X1,esk2_0)|esk1_0!=X1|~r2_hidden(esk8_3(X1,esk2_0,esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_28]), c_0_26]), c_0_27])])).
cnf(c_0_33, plain, (v1_funct_2(X1,X2,X3)|r2_hidden(esk8_3(X2,X3,X1),X2)|k1_relat_1(X1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_34, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|r2_hidden(esk8_3(X2,X3,X1),X2)|k1_relat_1(X1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_35, negated_conjecture, (~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))|~v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_27])])).
cnf(c_0_36, negated_conjecture, (v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_28]), c_0_26]), c_0_27])])).
cnf(c_0_37, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|r2_hidden(esk8_3(X1,X2,esk3_0),X1)|esk1_0!=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_28]), c_0_26]), c_0_27])])).
cnf(c_0_38, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r2_hidden(k1_funct_1(X1,esk8_3(X2,X3,X1)),X3)|k1_relat_1(X1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_39, negated_conjecture, (~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_36])])).
cnf(c_0_40, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,X1)))|r2_hidden(esk8_3(esk1_0,X1,esk3_0),esk1_0)), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_41, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,esk2_0)))|esk1_0!=X1|~r2_hidden(esk8_3(X1,esk2_0,esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_30]), c_0_28]), c_0_26]), c_0_27])])).
cnf(c_0_42, negated_conjecture, (r2_hidden(esk8_3(esk1_0,esk2_0,esk3_0),esk1_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_43, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_41]), c_0_42])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 44
# Proof object clause steps            : 33
# Proof object formula steps           : 11
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 30
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 22
# Processed clauses                    : 68
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 68
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 75
# ...of the previous two non-trivial   : 73
# Contextual simplify-reflections      : 0
# Paramodulations                      : 67
# Factorizations                       : 0
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 43
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 33
# Current number of unprocessed clauses: 49
# ...number of literals in the above   : 181
# Current number of archived formulas  : 0
# Current number of archived clauses   : 24
# Clause-clause subsumption calls (NU) : 98
# Rec. Clause-clause subsumption calls : 34
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 86
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2979

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.004 s
# Total time               : 0.011 s
# Maximum resident set size: 28648 pages
