# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t36_yellow_0, conjecture, ![X1]:((v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))=>![X2, X3]:((r1_yellow_0(X1,X2)&(r1_yellow_0(X1,X3)&r1_yellow_0(X1,k2_xboole_0(X2,X3))))=>k1_yellow_0(X1,k2_xboole_0(X2,X3))=k10_lattice3(X1,k1_yellow_0(X1,X2),k1_yellow_0(X1,X3)))), file('data/MPTP/proofs/38ee8eee28264d1f9f2ae69cb83a41e4.E_input', t36_yellow_0)).
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/38ee8eee28264d1f9f2ae69cb83a41e4.E_input', d9_yellow_0)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/38ee8eee28264d1f9f2ae69cb83a41e4.E_input', dt_k1_yellow_0)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/38ee8eee28264d1f9f2ae69cb83a41e4.E_input', t98_xboole_1)).
fof(t9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(r1_tarski(X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X3,X4)=>r1_lattice3(X1,X2,X4))&(r2_lattice3(X1,X3,X4)=>r2_lattice3(X1,X2,X4)))))), file('data/MPTP/proofs/38ee8eee28264d1f9f2ae69cb83a41e4.E_input', t9_yellow_0)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/38ee8eee28264d1f9f2ae69cb83a41e4.E_input', t7_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/38ee8eee28264d1f9f2ae69cb83a41e4.E_input', commutativity_k2_xboole_0)).
fof(d7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X5)=>r1_orders_2(X1,X4,X5))))=>X4=X3))))))), file('data/MPTP/proofs/38ee8eee28264d1f9f2ae69cb83a41e4.E_input', d7_yellow_0)).
fof(t18_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))&((r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))=>(X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3))))))))), file('data/MPTP/proofs/38ee8eee28264d1f9f2ae69cb83a41e4.E_input', t18_yellow_0)).
fof(t34_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:((r1_tarski(X2,X3)&(r1_yellow_0(X1,X2)&r1_yellow_0(X1,X3)))=>r1_orders_2(X1,k1_yellow_0(X1,X2),k1_yellow_0(X1,X3)))), file('data/MPTP/proofs/38ee8eee28264d1f9f2ae69cb83a41e4.E_input', t34_yellow_0)).
fof(t4_yellow_0, axiom, ![X1]:((v4_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)=>![X4]:((r1_lattice3(X1,X4,X3)=>r1_lattice3(X1,X4,X2))&(r2_lattice3(X1,X4,X2)=>r2_lattice3(X1,X4,X3))))))), file('data/MPTP/proofs/38ee8eee28264d1f9f2ae69cb83a41e4.E_input', t4_yellow_0)).
fof(t10_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3, X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((r1_lattice3(X1,X2,X4)&r1_lattice3(X1,X3,X4))=>r1_lattice3(X1,k2_xboole_0(X2,X3),X4))&((r2_lattice3(X1,X2,X4)&r2_lattice3(X1,X3,X4))=>r2_lattice3(X1,k2_xboole_0(X2,X3),X4))))), file('data/MPTP/proofs/38ee8eee28264d1f9f2ae69cb83a41e4.E_input', t10_yellow_0)).
fof(c_0_12, negated_conjecture, ~(![X1]:((v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))=>![X2, X3]:((r1_yellow_0(X1,X2)&(r1_yellow_0(X1,X3)&r1_yellow_0(X1,k2_xboole_0(X2,X3))))=>k1_yellow_0(X1,k2_xboole_0(X2,X3))=k10_lattice3(X1,k1_yellow_0(X1,X2),k1_yellow_0(X1,X3))))), inference(assume_negation,[status(cth)],[t36_yellow_0])).
fof(c_0_13, plain, ![X9, X10, X11, X12]:(((r2_lattice3(X9,X10,X11)|X11!=k1_yellow_0(X9,X10)|~r1_yellow_0(X9,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~l1_orders_2(X9))&(~m1_subset_1(X12,u1_struct_0(X9))|(~r2_lattice3(X9,X10,X12)|r1_orders_2(X9,X11,X12))|X11!=k1_yellow_0(X9,X10)|~r1_yellow_0(X9,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~l1_orders_2(X9)))&((m1_subset_1(esk4_3(X9,X10,X11),u1_struct_0(X9))|~r2_lattice3(X9,X10,X11)|X11=k1_yellow_0(X9,X10)|~r1_yellow_0(X9,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~l1_orders_2(X9))&((r2_lattice3(X9,X10,esk4_3(X9,X10,X11))|~r2_lattice3(X9,X10,X11)|X11=k1_yellow_0(X9,X10)|~r1_yellow_0(X9,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~l1_orders_2(X9))&(~r1_orders_2(X9,X11,esk4_3(X9,X10,X11))|~r2_lattice3(X9,X10,X11)|X11=k1_yellow_0(X9,X10)|~r1_yellow_0(X9,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~l1_orders_2(X9))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
fof(c_0_14, plain, ![X14, X15]:(~l1_orders_2(X14)|m1_subset_1(k1_yellow_0(X14,X15),u1_struct_0(X14))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
fof(c_0_15, negated_conjecture, ((v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))&((r1_yellow_0(esk1_0,esk2_0)&(r1_yellow_0(esk1_0,esk3_0)&r1_yellow_0(esk1_0,k2_xboole_0(esk2_0,esk3_0))))&k1_yellow_0(esk1_0,k2_xboole_0(esk2_0,esk3_0))!=k10_lattice3(esk1_0,k1_yellow_0(esk1_0,esk2_0),k1_yellow_0(esk1_0,esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
fof(c_0_16, plain, ![X16, X17]:k2_xboole_0(X16,X17)=k5_xboole_0(X16,k4_xboole_0(X17,X16)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_17, plain, ![X18, X19, X20, X21]:((~r1_lattice3(X18,X20,X21)|r1_lattice3(X18,X19,X21)|~m1_subset_1(X21,u1_struct_0(X18))|~r1_tarski(X19,X20)|~l1_orders_2(X18))&(~r2_lattice3(X18,X20,X21)|r2_lattice3(X18,X19,X21)|~m1_subset_1(X21,u1_struct_0(X18))|~r1_tarski(X19,X20)|~l1_orders_2(X18))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_yellow_0])])])])).
cnf(c_0_18, plain, (r2_lattice3(X1,X2,X3)|X3!=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, negated_conjecture, (r1_yellow_0(esk1_0,k2_xboole_0(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, plain, ![X22, X23]:r1_tarski(X22,k2_xboole_0(X22,X23)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_23, plain, ![X24, X25]:k2_xboole_0(X24,X25)=k2_xboole_0(X25,X24), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_24, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))&((r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))=>(X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3))))))), introduced(definition)).
fof(c_0_25, plain, ![X26, X27, X29, X30, X32, X33, X36]:(((m1_subset_1(esk5_2(X26,X27),u1_struct_0(X26))|~r1_yellow_0(X26,X27)|~l1_orders_2(X26))&((r2_lattice3(X26,X27,esk5_2(X26,X27))|~r1_yellow_0(X26,X27)|~l1_orders_2(X26))&((~m1_subset_1(X29,u1_struct_0(X26))|(~r2_lattice3(X26,X27,X29)|r1_orders_2(X26,esk5_2(X26,X27),X29))|~r1_yellow_0(X26,X27)|~l1_orders_2(X26))&((m1_subset_1(esk6_3(X26,X27,X30),u1_struct_0(X26))|~r2_lattice3(X26,X27,X30)|X30=esk5_2(X26,X27)|~m1_subset_1(X30,u1_struct_0(X26))|~r1_yellow_0(X26,X27)|~l1_orders_2(X26))&((r2_lattice3(X26,X27,esk6_3(X26,X27,X30))|~r2_lattice3(X26,X27,X30)|X30=esk5_2(X26,X27)|~m1_subset_1(X30,u1_struct_0(X26))|~r1_yellow_0(X26,X27)|~l1_orders_2(X26))&(~r1_orders_2(X26,X30,esk6_3(X26,X27,X30))|~r2_lattice3(X26,X27,X30)|X30=esk5_2(X26,X27)|~m1_subset_1(X30,u1_struct_0(X26))|~r1_yellow_0(X26,X27)|~l1_orders_2(X26)))))))&(((m1_subset_1(esk8_3(X26,X32,X33),u1_struct_0(X26))|m1_subset_1(esk7_3(X26,X32,X33),u1_struct_0(X26))|~r2_lattice3(X26,X32,X33)|~m1_subset_1(X33,u1_struct_0(X26))|r1_yellow_0(X26,X32)|~l1_orders_2(X26))&(((r2_lattice3(X26,X32,esk8_3(X26,X32,X33))|m1_subset_1(esk7_3(X26,X32,X33),u1_struct_0(X26))|~r2_lattice3(X26,X32,X33)|~m1_subset_1(X33,u1_struct_0(X26))|r1_yellow_0(X26,X32)|~l1_orders_2(X26))&(~m1_subset_1(X36,u1_struct_0(X26))|(~r2_lattice3(X26,X32,X36)|r1_orders_2(X26,esk8_3(X26,X32,X33),X36))|m1_subset_1(esk7_3(X26,X32,X33),u1_struct_0(X26))|~r2_lattice3(X26,X32,X33)|~m1_subset_1(X33,u1_struct_0(X26))|r1_yellow_0(X26,X32)|~l1_orders_2(X26)))&(esk8_3(X26,X32,X33)!=X33|m1_subset_1(esk7_3(X26,X32,X33),u1_struct_0(X26))|~r2_lattice3(X26,X32,X33)|~m1_subset_1(X33,u1_struct_0(X26))|r1_yellow_0(X26,X32)|~l1_orders_2(X26))))&(((m1_subset_1(esk8_3(X26,X32,X33),u1_struct_0(X26))|r2_lattice3(X26,X32,esk7_3(X26,X32,X33))|~r2_lattice3(X26,X32,X33)|~m1_subset_1(X33,u1_struct_0(X26))|r1_yellow_0(X26,X32)|~l1_orders_2(X26))&(((r2_lattice3(X26,X32,esk8_3(X26,X32,X33))|r2_lattice3(X26,X32,esk7_3(X26,X32,X33))|~r2_lattice3(X26,X32,X33)|~m1_subset_1(X33,u1_struct_0(X26))|r1_yellow_0(X26,X32)|~l1_orders_2(X26))&(~m1_subset_1(X36,u1_struct_0(X26))|(~r2_lattice3(X26,X32,X36)|r1_orders_2(X26,esk8_3(X26,X32,X33),X36))|r2_lattice3(X26,X32,esk7_3(X26,X32,X33))|~r2_lattice3(X26,X32,X33)|~m1_subset_1(X33,u1_struct_0(X26))|r1_yellow_0(X26,X32)|~l1_orders_2(X26)))&(esk8_3(X26,X32,X33)!=X33|r2_lattice3(X26,X32,esk7_3(X26,X32,X33))|~r2_lattice3(X26,X32,X33)|~m1_subset_1(X33,u1_struct_0(X26))|r1_yellow_0(X26,X32)|~l1_orders_2(X26))))&((m1_subset_1(esk8_3(X26,X32,X33),u1_struct_0(X26))|~r1_orders_2(X26,X33,esk7_3(X26,X32,X33))|~r2_lattice3(X26,X32,X33)|~m1_subset_1(X33,u1_struct_0(X26))|r1_yellow_0(X26,X32)|~l1_orders_2(X26))&(((r2_lattice3(X26,X32,esk8_3(X26,X32,X33))|~r1_orders_2(X26,X33,esk7_3(X26,X32,X33))|~r2_lattice3(X26,X32,X33)|~m1_subset_1(X33,u1_struct_0(X26))|r1_yellow_0(X26,X32)|~l1_orders_2(X26))&(~m1_subset_1(X36,u1_struct_0(X26))|(~r2_lattice3(X26,X32,X36)|r1_orders_2(X26,esk8_3(X26,X32,X33),X36))|~r1_orders_2(X26,X33,esk7_3(X26,X32,X33))|~r2_lattice3(X26,X32,X33)|~m1_subset_1(X33,u1_struct_0(X26))|r1_yellow_0(X26,X32)|~l1_orders_2(X26)))&(esk8_3(X26,X32,X33)!=X33|~r1_orders_2(X26,X33,esk7_3(X26,X32,X33))|~r2_lattice3(X26,X32,X33)|~m1_subset_1(X33,u1_struct_0(X26))|r1_yellow_0(X26,X32)|~l1_orders_2(X26))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_yellow_0])])])])])])).
cnf(c_0_26, plain, (r2_lattice3(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_tarski(X4,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27, plain, (r2_lattice3(X1,X2,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_18]), c_0_19])).
cnf(c_0_28, negated_conjecture, (r1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_29, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_30, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_32, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>epred1_3(X1,X2,X3)))), inference(apply_def,[status(thm)],[t18_yellow_0, c_0_24])).
cnf(c_0_33, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|X3=esk5_2(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, negated_conjecture, (r1_yellow_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_35, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))&((r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))=>(X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3))))))), inference(split_equiv,[status(thm)],[c_0_24])).
cnf(c_0_36, plain, (r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~r1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_37, plain, (r2_lattice3(X1,X2,k1_yellow_0(X1,X3))|~r1_tarski(X2,X4)|~r2_lattice3(X1,X4,k1_yellow_0(X1,X3))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_26, c_0_19])).
cnf(c_0_38, negated_conjecture, (r2_lattice3(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)),k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])])).
cnf(c_0_39, plain, (r1_tarski(X1,k5_xboole_0(X1,k4_xboole_0(X2,X1)))), inference(rw,[status(thm)],[c_0_30, c_0_21])).
cnf(c_0_40, plain, (k5_xboole_0(X1,k4_xboole_0(X2,X1))=k5_xboole_0(X2,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_21]), c_0_21])).
fof(c_0_41, plain, ![X37, X38, X39]:(~v5_orders_2(X37)|~l1_orders_2(X37)|(~m1_subset_1(X38,u1_struct_0(X37))|(~m1_subset_1(X39,u1_struct_0(X37))|epred1_3(X37,X38,X39)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_32])])])).
cnf(c_0_42, plain, (k1_yellow_0(X1,X2)=esk5_2(X1,X3)|m1_subset_1(esk6_3(X1,X3,k1_yellow_0(X1,X2)),u1_struct_0(X1))|~r2_lattice3(X1,X3,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X3)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_33, c_0_19])).
cnf(c_0_43, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_34]), c_0_29])])).
cnf(c_0_44, plain, (r2_lattice3(X1,X2,esk6_3(X1,X2,X3))|X3=esk5_2(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_45, plain, (X2=esk5_2(X1,X3)|~r1_orders_2(X1,X2,esk6_3(X1,X3,X2))|~r2_lattice3(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_yellow_0(X1,X3)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_46, plain, ![X51, X52, X53, X54, X55]:(((r1_orders_2(X53,X52,X54)|(X54!=k10_lattice3(X53,X52,X51)|~r1_yellow_0(X53,k2_tarski(X52,X51)))|~m1_subset_1(X54,u1_struct_0(X53))|~epred1_3(X53,X52,X51))&((r1_orders_2(X53,X51,X54)|(X54!=k10_lattice3(X53,X52,X51)|~r1_yellow_0(X53,k2_tarski(X52,X51)))|~m1_subset_1(X54,u1_struct_0(X53))|~epred1_3(X53,X52,X51))&(~m1_subset_1(X55,u1_struct_0(X53))|(~r1_orders_2(X53,X52,X55)|~r1_orders_2(X53,X51,X55)|r1_orders_2(X53,X54,X55))|(X54!=k10_lattice3(X53,X52,X51)|~r1_yellow_0(X53,k2_tarski(X52,X51)))|~m1_subset_1(X54,u1_struct_0(X53))|~epred1_3(X53,X52,X51))))&(((X54=k10_lattice3(X53,X52,X51)|(m1_subset_1(esk9_4(X51,X52,X53,X54),u1_struct_0(X53))|~r1_orders_2(X53,X51,X54)|~r1_orders_2(X53,X52,X54))|~m1_subset_1(X54,u1_struct_0(X53))|~epred1_3(X53,X52,X51))&(r1_yellow_0(X53,k2_tarski(X52,X51))|(m1_subset_1(esk9_4(X51,X52,X53,X54),u1_struct_0(X53))|~r1_orders_2(X53,X51,X54)|~r1_orders_2(X53,X52,X54))|~m1_subset_1(X54,u1_struct_0(X53))|~epred1_3(X53,X52,X51)))&((((X54=k10_lattice3(X53,X52,X51)|(r1_orders_2(X53,X52,esk9_4(X51,X52,X53,X54))|~r1_orders_2(X53,X51,X54)|~r1_orders_2(X53,X52,X54))|~m1_subset_1(X54,u1_struct_0(X53))|~epred1_3(X53,X52,X51))&(r1_yellow_0(X53,k2_tarski(X52,X51))|(r1_orders_2(X53,X52,esk9_4(X51,X52,X53,X54))|~r1_orders_2(X53,X51,X54)|~r1_orders_2(X53,X52,X54))|~m1_subset_1(X54,u1_struct_0(X53))|~epred1_3(X53,X52,X51)))&((X54=k10_lattice3(X53,X52,X51)|(r1_orders_2(X53,X51,esk9_4(X51,X52,X53,X54))|~r1_orders_2(X53,X51,X54)|~r1_orders_2(X53,X52,X54))|~m1_subset_1(X54,u1_struct_0(X53))|~epred1_3(X53,X52,X51))&(r1_yellow_0(X53,k2_tarski(X52,X51))|(r1_orders_2(X53,X51,esk9_4(X51,X52,X53,X54))|~r1_orders_2(X53,X51,X54)|~r1_orders_2(X53,X52,X54))|~m1_subset_1(X54,u1_struct_0(X53))|~epred1_3(X53,X52,X51))))&((X54=k10_lattice3(X53,X52,X51)|(~r1_orders_2(X53,X54,esk9_4(X51,X52,X53,X54))|~r1_orders_2(X53,X51,X54)|~r1_orders_2(X53,X52,X54))|~m1_subset_1(X54,u1_struct_0(X53))|~epred1_3(X53,X52,X51))&(r1_yellow_0(X53,k2_tarski(X52,X51))|(~r1_orders_2(X53,X54,esk9_4(X51,X52,X53,X54))|~r1_orders_2(X53,X51,X54)|~r1_orders_2(X53,X52,X54))|~m1_subset_1(X54,u1_struct_0(X53))|~epred1_3(X53,X52,X51)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_35])])])])])).
cnf(c_0_47, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_36]), c_0_19])).
cnf(c_0_48, negated_conjecture, (r2_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0))))|~r1_tarski(X1,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_29])])).
cnf(c_0_49, plain, (r1_tarski(X1,k5_xboole_0(X2,k4_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_50, plain, (epred1_3(X1,X2,X3)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_51, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_52, negated_conjecture, (esk5_2(esk1_0,esk2_0)=k1_yellow_0(esk1_0,esk2_0)|m1_subset_1(esk6_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_34]), c_0_29])])).
cnf(c_0_53, plain, (k1_yellow_0(X1,X2)=esk5_2(X1,X3)|r2_lattice3(X1,X3,esk6_3(X1,X3,k1_yellow_0(X1,X2)))|~r2_lattice3(X1,X3,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X3)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_44, c_0_19])).
cnf(c_0_54, plain, (k1_yellow_0(X1,X2)=esk5_2(X1,X3)|~r1_orders_2(X1,k1_yellow_0(X1,X2),esk6_3(X1,X3,k1_yellow_0(X1,X2)))|~r2_lattice3(X1,X3,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X3)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_45, c_0_19])).
fof(c_0_55, plain, ![X40, X41, X42]:(~l1_orders_2(X40)|(~r1_tarski(X41,X42)|(~r1_yellow_0(X40,X41)|~r1_yellow_0(X40,X42))|r1_orders_2(X40,k1_yellow_0(X40,X41),k1_yellow_0(X40,X42)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t34_yellow_0])])])).
cnf(c_0_56, plain, (X1=k10_lattice3(X2,X3,X4)|m1_subset_1(esk9_4(X4,X3,X2,X1),u1_struct_0(X2))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_57, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),k1_yellow_0(X1,X3))|~r2_lattice3(X1,X2,k1_yellow_0(X1,X3))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_47, c_0_19])).
cnf(c_0_58, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0))))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_59, negated_conjecture, (r1_yellow_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_60, plain, (epred1_3(X1,X2,k1_yellow_0(X1,X3))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_50, c_0_19])).
cnf(c_0_61, negated_conjecture, (m1_subset_1(esk5_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_34]), c_0_29])])).
cnf(c_0_62, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_63, negated_conjecture, (esk5_2(esk1_0,esk2_0)=k1_yellow_0(esk1_0,esk2_0)|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),esk6_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk2_0)))|~r2_lattice3(esk1_0,X1,esk6_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk2_0)))|~r1_yellow_0(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_52]), c_0_29])])).
cnf(c_0_64, negated_conjecture, (esk5_2(esk1_0,esk2_0)=k1_yellow_0(esk1_0,esk2_0)|r2_lattice3(esk1_0,esk2_0,esk6_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_43]), c_0_34]), c_0_29])])).
cnf(c_0_65, negated_conjecture, (esk5_2(esk1_0,esk2_0)=k1_yellow_0(esk1_0,esk2_0)|~r1_orders_2(esk1_0,k1_yellow_0(esk1_0,esk2_0),esk6_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_43]), c_0_34]), c_0_29])])).
cnf(c_0_66, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),k1_yellow_0(X1,X3))|~l1_orders_2(X1)|~r1_tarski(X2,X3)|~r1_yellow_0(X1,X2)|~r1_yellow_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_67, plain, (k1_yellow_0(X1,X2)=k10_lattice3(X1,X3,X4)|m1_subset_1(esk9_4(X4,X3,X1,k1_yellow_0(X1,X2)),u1_struct_0(X1))|~epred1_3(X1,X3,X4)|~r1_orders_2(X1,X4,k1_yellow_0(X1,X2))|~r1_orders_2(X1,X3,k1_yellow_0(X1,X2))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_56, c_0_19])).
cnf(c_0_68, negated_conjecture, (r1_orders_2(esk1_0,k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59]), c_0_29])])).
cnf(c_0_69, negated_conjecture, (epred1_3(esk1_0,esk5_2(esk1_0,esk2_0),k1_yellow_0(esk1_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_29]), c_0_62])])).
cnf(c_0_70, negated_conjecture, (esk5_2(esk1_0,esk2_0)=k1_yellow_0(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_34]), c_0_64]), c_0_65])).
cnf(c_0_71, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),k1_yellow_0(X1,k5_xboole_0(X2,k4_xboole_0(X3,X2))))|~r1_yellow_0(X1,k5_xboole_0(X2,k4_xboole_0(X3,X2)))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_66, c_0_39])).
cnf(c_0_72, negated_conjecture, (k1_yellow_0(esk1_0,k2_xboole_0(esk2_0,esk3_0))!=k10_lattice3(esk1_0,k1_yellow_0(esk1_0,esk2_0),k1_yellow_0(esk1_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_73, plain, ![X43, X44, X45, X46]:((~r1_lattice3(X43,X46,X45)|r1_lattice3(X43,X46,X44)|~r1_orders_2(X43,X44,X45)|~m1_subset_1(X45,u1_struct_0(X43))|~m1_subset_1(X44,u1_struct_0(X43))|(~v4_orders_2(X43)|~l1_orders_2(X43)))&(~r2_lattice3(X43,X46,X44)|r2_lattice3(X43,X46,X45)|~r1_orders_2(X43,X44,X45)|~m1_subset_1(X45,u1_struct_0(X43))|~m1_subset_1(X44,u1_struct_0(X43))|(~v4_orders_2(X43)|~l1_orders_2(X43)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_yellow_0])])])])).
cnf(c_0_74, negated_conjecture, (k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))=k10_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0))|m1_subset_1(esk9_4(k1_yellow_0(esk1_0,esk3_0),X1,esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))),u1_struct_0(esk1_0))|~epred1_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0))|~r1_orders_2(esk1_0,X1,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_29])])).
cnf(c_0_75, negated_conjecture, (epred1_3(esk1_0,k1_yellow_0(esk1_0,esk2_0),k1_yellow_0(esk1_0,X1))), inference(rw,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_76, negated_conjecture, (r1_orders_2(esk1_0,k1_yellow_0(esk1_0,esk2_0),k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_28]), c_0_34]), c_0_29])])).
cnf(c_0_77, negated_conjecture, (k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))!=k10_lattice3(esk1_0,k1_yellow_0(esk1_0,esk2_0),k1_yellow_0(esk1_0,esk3_0))), inference(rw,[status(thm)],[c_0_72, c_0_21])).
cnf(c_0_78, plain, (r2_lattice3(X1,X2,X4)|~r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,X3,X4)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_79, negated_conjecture, (m1_subset_1(esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76])]), c_0_77])).
cnf(c_0_80, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_81, plain, (X1=k10_lattice3(X2,X3,X4)|r1_orders_2(X2,X4,esk9_4(X4,X3,X2,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_82, plain, ![X47, X48, X49, X50]:((~r1_lattice3(X47,X48,X50)|~r1_lattice3(X47,X49,X50)|r1_lattice3(X47,k2_xboole_0(X48,X49),X50)|~m1_subset_1(X50,u1_struct_0(X47))|~l1_orders_2(X47))&(~r2_lattice3(X47,X48,X50)|~r2_lattice3(X47,X49,X50)|r2_lattice3(X47,k2_xboole_0(X48,X49),X50)|~m1_subset_1(X50,u1_struct_0(X47))|~l1_orders_2(X47))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_yellow_0])])])])).
cnf(c_0_83, negated_conjecture, (r2_lattice3(esk1_0,X1,esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))|~r1_orders_2(esk1_0,X2,esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))|~r2_lattice3(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_29]), c_0_80])])).
cnf(c_0_84, plain, (k1_yellow_0(X1,X2)=k10_lattice3(X1,X3,X4)|r1_orders_2(X1,X4,esk9_4(X4,X3,X1,k1_yellow_0(X1,X2)))|~epred1_3(X1,X3,X4)|~r1_orders_2(X1,X4,k1_yellow_0(X1,X2))|~r1_orders_2(X1,X3,k1_yellow_0(X1,X2))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_81, c_0_19])).
cnf(c_0_85, plain, (X1=k10_lattice3(X2,X3,X4)|r1_orders_2(X2,X3,esk9_4(X4,X3,X2,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_86, plain, (X1=k10_lattice3(X2,X3,X4)|~r1_orders_2(X2,X1,esk9_4(X4,X3,X2,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_87, plain, (r2_lattice3(X1,k2_xboole_0(X2,X4),X3)|~r2_lattice3(X1,X2,X3)|~r2_lattice3(X1,X4,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_88, negated_conjecture, (r2_lattice3(esk1_0,X1,esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))|~r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X2),esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))|~r2_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_19]), c_0_29])])).
cnf(c_0_89, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_59]), c_0_29])])).
cnf(c_0_90, negated_conjecture, (k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))=k10_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0))|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,esk3_0),esk9_4(k1_yellow_0(esk1_0,esk3_0),X1,esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))|~epred1_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0))|~r1_orders_2(esk1_0,X1,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_68]), c_0_29])])).
cnf(c_0_91, plain, (k1_yellow_0(X1,X2)=k10_lattice3(X1,X3,X4)|r1_orders_2(X1,X3,esk9_4(X4,X3,X1,k1_yellow_0(X1,X2)))|~epred1_3(X1,X3,X4)|~r1_orders_2(X1,X4,k1_yellow_0(X1,X2))|~r1_orders_2(X1,X3,k1_yellow_0(X1,X2))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_85, c_0_19])).
cnf(c_0_92, plain, (k1_yellow_0(X1,X2)=k10_lattice3(X1,X3,X4)|~epred1_3(X1,X3,X4)|~r1_orders_2(X1,k1_yellow_0(X1,X2),esk9_4(X4,X3,X1,k1_yellow_0(X1,X2)))|~r1_orders_2(X1,X4,k1_yellow_0(X1,X2))|~r1_orders_2(X1,X3,k1_yellow_0(X1,X2))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_86, c_0_19])).
cnf(c_0_93, plain, (r2_lattice3(X1,k5_xboole_0(X2,k4_xboole_0(X4,X2)),X3)|~l1_orders_2(X1)|~r2_lattice3(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(rw,[status(thm)],[c_0_87, c_0_21])).
cnf(c_0_94, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))|~r1_orders_2(esk1_0,k1_yellow_0(esk1_0,esk3_0),esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_95, negated_conjecture, (r1_orders_2(esk1_0,k1_yellow_0(esk1_0,esk3_0),esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_75]), c_0_76])]), c_0_77])).
cnf(c_0_96, negated_conjecture, (k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))=k10_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0))|r1_orders_2(esk1_0,X1,esk9_4(k1_yellow_0(esk1_0,esk3_0),X1,esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))|~epred1_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0))|~r1_orders_2(esk1_0,X1,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_68]), c_0_29])])).
cnf(c_0_97, negated_conjecture, (k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))=k10_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0))|~epred1_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0))|~r1_orders_2(esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0))),esk9_4(k1_yellow_0(esk1_0,esk3_0),X1,esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))|~r1_orders_2(esk1_0,X1,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_68]), c_0_29])])).
cnf(c_0_98, negated_conjecture, (r2_lattice3(esk1_0,k5_xboole_0(X1,k4_xboole_0(X2,X1)),esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))|~r2_lattice3(esk1_0,X2,esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))|~r2_lattice3(esk1_0,X1,esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_79]), c_0_29])])).
cnf(c_0_99, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_94, c_0_95])])).
cnf(c_0_100, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))|~r1_orders_2(esk1_0,k1_yellow_0(esk1_0,esk2_0),esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))), inference(spm,[status(thm)],[c_0_88, c_0_43])).
cnf(c_0_101, negated_conjecture, (r1_orders_2(esk1_0,k1_yellow_0(esk1_0,esk2_0),esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_75]), c_0_76])]), c_0_77])).
cnf(c_0_102, negated_conjecture, (r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))|~r2_lattice3(esk1_0,X1,esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))|~r1_yellow_0(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_79]), c_0_29])])).
cnf(c_0_103, negated_conjecture, (~r1_orders_2(esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0))),esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_75]), c_0_76])]), c_0_77])).
cnf(c_0_104, negated_conjecture, (r2_lattice3(esk1_0,k5_xboole_0(X1,k4_xboole_0(esk3_0,X1)),esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))|~r2_lattice3(esk1_0,X1,esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_105, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_100, c_0_101])])).
cnf(c_0_106, negated_conjecture, (~r2_lattice3(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)),esk9_4(k1_yellow_0(esk1_0,esk3_0),k1_yellow_0(esk1_0,esk2_0),esk1_0,k1_yellow_0(esk1_0,k5_xboole_0(esk2_0,k4_xboole_0(esk3_0,esk2_0)))))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_28]), c_0_103])).
cnf(c_0_107, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_106]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 108
# Proof object clause steps            : 79
# Proof object formula steps           : 29
# Proof object conjectures             : 46
# Proof object clause conjectures      : 43
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 43
# Proof object simplifying inferences  : 78
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 53
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 52
# Processed clauses                    : 3204
# ...of these trivial                  : 54
# ...subsumed                          : 1287
# ...remaining for further processing  : 1863
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 128
# Backward-rewritten                   : 371
# Generated clauses                    : 18608
# ...of the previous two non-trivial   : 17443
# Contextual simplify-reflections      : 23
# Paramodulations                      : 18603
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1307
#    Positive orientable unit clauses  : 240
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 1063
# Current number of unprocessed clauses: 14009
# ...number of literals in the above   : 53127
# Current number of archived formulas  : 0
# Current number of archived clauses   : 552
# Clause-clause subsumption calls (NU) : 79255
# Rec. Clause-clause subsumption calls : 63925
# Non-unit clause-clause subsumptions  : 1423
# Unit Clause-clause subsumption calls : 12660
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5808
# BW rewrite match successes           : 47
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1219035

# -------------------------------------------------
# User time                : 0.385 s
# System time              : 0.007 s
# Total time               : 0.392 s
# Maximum resident set size: 28788 pages
