# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d18_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k10_yellow_6(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>![X5]:(m1_connsp_2(X5,X1,X4)=>r1_waybel_0(X1,X2,X5)))))))), file('data/MPTP/proofs/38f734f43a69409cbce4b4c18300fe84.E_input', d18_yellow_6)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/38f734f43a69409cbce4b4c18300fe84.E_input', t4_subset)).
fof(dt_k10_yellow_6, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&(~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1)))))=>m1_subset_1(k10_yellow_6(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/38f734f43a69409cbce4b4c18300fe84.E_input', dt_k10_yellow_6)).
fof(t24_waybel_9, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((~(v2_struct_0(X3))&(v4_orders_2(X3)&(v7_waybel_0(X3)&l1_waybel_0(X3,X1))))=>(r2_hidden(X2,k10_yellow_6(X1,X3))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(X4=k2_relset_1(u1_struct_0(X3),u1_struct_0(X1),u1_waybel_0(X1,X3))=>r2_hidden(X2,k2_pre_topc(X1,X4)))))))), file('data/MPTP/proofs/38f734f43a69409cbce4b4c18300fe84.E_input', t24_waybel_9)).
fof(t5_connsp_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((v3_pre_topc(X2,X1)&r2_hidden(X3,X2))=>m1_connsp_2(X2,X1,X3))))), file('data/MPTP/proofs/38f734f43a69409cbce4b4c18300fe84.E_input', t5_connsp_2)).
fof(t39_tops_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k2_pre_topc(X1,X2))<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>~((v3_pre_topc(X4,X1)&(r2_hidden(X3,X4)&r1_xboole_0(X2,X4))))))))), file('data/MPTP/proofs/38f734f43a69409cbce4b4c18300fe84.E_input', t39_tops_1)).
fof(t26_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3, X4]:~((r1_waybel_0(X1,X2,X3)&(r1_waybel_0(X1,X2,X4)&r1_xboole_0(X3,X4)))))), file('data/MPTP/proofs/38f734f43a69409cbce4b4c18300fe84.E_input', t26_yellow_6)).
fof(t8_waybel_9, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>r1_waybel_0(X1,X2,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2))))), file('data/MPTP/proofs/38f734f43a69409cbce4b4c18300fe84.E_input', t8_waybel_9)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/38f734f43a69409cbce4b4c18300fe84.E_input', dt_l1_pre_topc)).
fof(c_0_9, plain, ![X21, X22, X23, X24, X25, X29]:(((~r2_hidden(X24,X23)|(~m1_connsp_2(X25,X21,X24)|r1_waybel_0(X21,X22,X25))|~m1_subset_1(X24,u1_struct_0(X21))|X23!=k10_yellow_6(X21,X22)|~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(X21)))|(v2_struct_0(X22)|(~v4_orders_2(X22)|(~v7_waybel_0(X22)|~l1_waybel_0(X22,X21))))|(v2_struct_0(X21)|(~v2_pre_topc(X21)|~l1_pre_topc(X21))))&((m1_connsp_2(esk5_4(X21,X22,X23,X24),X21,X24)|r2_hidden(X24,X23)|~m1_subset_1(X24,u1_struct_0(X21))|X23!=k10_yellow_6(X21,X22)|~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(X21)))|(v2_struct_0(X22)|(~v4_orders_2(X22)|(~v7_waybel_0(X22)|~l1_waybel_0(X22,X21))))|(v2_struct_0(X21)|(~v2_pre_topc(X21)|~l1_pre_topc(X21))))&(~r1_waybel_0(X21,X22,esk5_4(X21,X22,X23,X24))|r2_hidden(X24,X23)|~m1_subset_1(X24,u1_struct_0(X21))|X23!=k10_yellow_6(X21,X22)|~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(X21)))|(v2_struct_0(X22)|(~v4_orders_2(X22)|(~v7_waybel_0(X22)|~l1_waybel_0(X22,X21))))|(v2_struct_0(X21)|(~v2_pre_topc(X21)|~l1_pre_topc(X21))))))&((m1_subset_1(esk6_3(X21,X22,X23),u1_struct_0(X21))|X23=k10_yellow_6(X21,X22)|~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(X21)))|(v2_struct_0(X22)|(~v4_orders_2(X22)|(~v7_waybel_0(X22)|~l1_waybel_0(X22,X21))))|(v2_struct_0(X21)|(~v2_pre_topc(X21)|~l1_pre_topc(X21))))&(((m1_connsp_2(esk7_3(X21,X22,X23),X21,esk6_3(X21,X22,X23))|~r2_hidden(esk6_3(X21,X22,X23),X23)|X23=k10_yellow_6(X21,X22)|~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(X21)))|(v2_struct_0(X22)|(~v4_orders_2(X22)|(~v7_waybel_0(X22)|~l1_waybel_0(X22,X21))))|(v2_struct_0(X21)|(~v2_pre_topc(X21)|~l1_pre_topc(X21))))&(~r1_waybel_0(X21,X22,esk7_3(X21,X22,X23))|~r2_hidden(esk6_3(X21,X22,X23),X23)|X23=k10_yellow_6(X21,X22)|~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(X21)))|(v2_struct_0(X22)|(~v4_orders_2(X22)|(~v7_waybel_0(X22)|~l1_waybel_0(X22,X21))))|(v2_struct_0(X21)|(~v2_pre_topc(X21)|~l1_pre_topc(X21)))))&(r2_hidden(esk6_3(X21,X22,X23),X23)|(~m1_connsp_2(X29,X21,esk6_3(X21,X22,X23))|r1_waybel_0(X21,X22,X29))|X23=k10_yellow_6(X21,X22)|~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(X21)))|(v2_struct_0(X22)|(~v4_orders_2(X22)|(~v7_waybel_0(X22)|~l1_waybel_0(X22,X21))))|(v2_struct_0(X21)|(~v2_pre_topc(X21)|~l1_pre_topc(X21))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d18_yellow_6])])])])])])).
fof(c_0_10, plain, ![X13, X14, X15]:(~r2_hidden(X13,X14)|~m1_subset_1(X14,k1_zfmisc_1(X15))|m1_subset_1(X13,X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_11, plain, (r1_waybel_0(X4,X5,X3)|v2_struct_0(X5)|v2_struct_0(X4)|~r2_hidden(X1,X2)|~m1_connsp_2(X3,X4,X1)|~m1_subset_1(X1,u1_struct_0(X4))|X2!=k10_yellow_6(X4,X5)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X4)))|~v4_orders_2(X5)|~v7_waybel_0(X5)|~l1_waybel_0(X5,X4)|~v2_pre_topc(X4)|~l1_pre_topc(X4)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_13, plain, ![X16, X17]:(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))|(v2_struct_0(X17)|(~v4_orders_2(X17)|(~v7_waybel_0(X17)|~l1_waybel_0(X17,X16))))|m1_subset_1(k10_yellow_6(X16,X17),k1_zfmisc_1(u1_struct_0(X16)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k10_yellow_6])])])).
fof(c_0_14, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((~(v2_struct_0(X3))&(v4_orders_2(X3)&(v7_waybel_0(X3)&l1_waybel_0(X3,X1))))=>(r2_hidden(X2,k10_yellow_6(X1,X3))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(X4=k2_relset_1(u1_struct_0(X3),u1_struct_0(X1),u1_waybel_0(X1,X3))=>r2_hidden(X2,k2_pre_topc(X1,X4))))))))), inference(assume_negation,[status(cth)],[t24_waybel_9])).
fof(c_0_15, plain, ![X10, X11, X12]:(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))|(~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|(~m1_subset_1(X12,u1_struct_0(X10))|(~v3_pre_topc(X11,X10)|~r2_hidden(X12,X11)|m1_connsp_2(X11,X10,X12))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t5_connsp_2])])])])).
cnf(c_0_16, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|X4!=k10_yellow_6(X1,X2)|~m1_connsp_2(X3,X1,X5)|~r2_hidden(X5,X4)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_17, plain, (v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(k10_yellow_6(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_18, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&((~v2_struct_0(esk3_0)&(v4_orders_2(esk3_0)&(v7_waybel_0(esk3_0)&l1_waybel_0(esk3_0,esk1_0))))&(r2_hidden(esk2_0,k10_yellow_6(esk1_0,esk3_0))&(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(esk4_0=k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0))&~r2_hidden(esk2_0,k2_pre_topc(esk1_0,esk4_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])).
cnf(c_0_19, plain, (v2_struct_0(X1)|m1_connsp_2(X2,X1,X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_pre_topc(X2,X1)|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_20, plain, ![X30, X31, X32, X33]:((~r2_hidden(X32,k2_pre_topc(X30,X31))|(~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X30)))|(~v3_pre_topc(X33,X30)|(~r2_hidden(X32,X33)|~r1_xboole_0(X31,X33))))|~m1_subset_1(X32,u1_struct_0(X30))|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))|(v2_struct_0(X30)|(~v2_pre_topc(X30)|~l1_pre_topc(X30))))&((m1_subset_1(esk8_3(X30,X31,X32),k1_zfmisc_1(u1_struct_0(X30)))|r2_hidden(X32,k2_pre_topc(X30,X31))|~m1_subset_1(X32,u1_struct_0(X30))|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))|(v2_struct_0(X30)|(~v2_pre_topc(X30)|~l1_pre_topc(X30))))&((v3_pre_topc(esk8_3(X30,X31,X32),X30)|r2_hidden(X32,k2_pre_topc(X30,X31))|~m1_subset_1(X32,u1_struct_0(X30))|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))|(v2_struct_0(X30)|(~v2_pre_topc(X30)|~l1_pre_topc(X30))))&((r2_hidden(X32,esk8_3(X30,X31,X32))|r2_hidden(X32,k2_pre_topc(X30,X31))|~m1_subset_1(X32,u1_struct_0(X30))|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))|(v2_struct_0(X30)|(~v2_pre_topc(X30)|~l1_pre_topc(X30))))&(r1_xboole_0(X31,esk8_3(X30,X31,X32))|r2_hidden(X32,k2_pre_topc(X30,X31))|~m1_subset_1(X32,u1_struct_0(X30))|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))|(v2_struct_0(X30)|(~v2_pre_topc(X30)|~l1_pre_topc(X30)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t39_tops_1])])])])])])).
cnf(c_0_21, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X2)|k10_yellow_6(X1,X4)!=k10_yellow_6(X1,X2)|~m1_connsp_2(X3,X1,X5)|~r2_hidden(X5,k10_yellow_6(X1,X4))|~l1_waybel_0(X2,X1)|~l1_waybel_0(X4,X1)|~v7_waybel_0(X2)|~v7_waybel_0(X4)|~v4_orders_2(X2)|~v4_orders_2(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_22, negated_conjecture, (r2_hidden(esk2_0,k10_yellow_6(esk1_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, negated_conjecture, (l1_waybel_0(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, negated_conjecture, (v7_waybel_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, negated_conjecture, (v4_orders_2(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_29, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_30, plain, (m1_connsp_2(X1,X2,X3)|v2_struct_0(X2)|~v3_pre_topc(X1,X2)|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[c_0_19, c_0_12])).
cnf(c_0_31, plain, (v3_pre_topc(esk8_3(X1,X2,X3),X1)|r2_hidden(X3,k2_pre_topc(X1,X2))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_32, plain, (m1_subset_1(esk8_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|r2_hidden(X3,k2_pre_topc(X1,X2))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_33, negated_conjecture, (r1_waybel_0(esk1_0,X1,X2)|v2_struct_0(X1)|k10_yellow_6(esk1_0,esk3_0)!=k10_yellow_6(esk1_0,X1)|~m1_connsp_2(X2,esk1_0,esk2_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27])]), c_0_28]), c_0_29])).
cnf(c_0_34, plain, (m1_connsp_2(esk8_3(X1,X2,X3),X1,X4)|r2_hidden(X3,k2_pre_topc(X1,X2))|v2_struct_0(X1)|~r2_hidden(X4,esk8_3(X1,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])).
fof(c_0_35, plain, ![X35, X36, X37, X38]:(v2_struct_0(X35)|~l1_struct_0(X35)|(v2_struct_0(X36)|(~v4_orders_2(X36)|(~v7_waybel_0(X36)|~l1_waybel_0(X36,X35)))|(~r1_waybel_0(X35,X36,X37)|(~r1_waybel_0(X35,X36,X38)|~r1_xboole_0(X37,X38))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t26_yellow_6])])])])).
cnf(c_0_36, negated_conjecture, (r1_waybel_0(esk1_0,X1,esk8_3(esk1_0,X2,X3))|r2_hidden(X3,k2_pre_topc(esk1_0,X2))|v2_struct_0(X1)|k10_yellow_6(esk1_0,esk3_0)!=k10_yellow_6(esk1_0,X1)|~r2_hidden(esk2_0,esk8_3(esk1_0,X2,X3))|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_26]), c_0_27])]), c_0_29])).
cnf(c_0_37, plain, (r2_hidden(X1,esk8_3(X2,X3,X1))|r2_hidden(X1,k2_pre_topc(X2,X3))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_38, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_39, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~l1_struct_0(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~r1_waybel_0(X1,X2,X3)|~r1_waybel_0(X1,X2,X4)|~r1_xboole_0(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_40, negated_conjecture, (r1_waybel_0(esk1_0,X1,esk8_3(esk1_0,X2,esk2_0))|r2_hidden(esk2_0,k2_pre_topc(esk1_0,X2))|v2_struct_0(X1)|k10_yellow_6(esk1_0,esk3_0)!=k10_yellow_6(esk1_0,X1)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_26]), c_0_27])]), c_0_29])).
fof(c_0_41, plain, ![X19, X20]:(v2_struct_0(X19)|~l1_struct_0(X19)|(v2_struct_0(X20)|~l1_waybel_0(X20,X19)|r1_waybel_0(X19,X20,k2_relset_1(u1_struct_0(X20),u1_struct_0(X19),u1_waybel_0(X19,X20))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_waybel_9])])])])).
cnf(c_0_42, negated_conjecture, (r2_hidden(esk2_0,k2_pre_topc(esk1_0,X1))|v2_struct_0(X2)|k10_yellow_6(esk1_0,esk3_0)!=k10_yellow_6(esk1_0,X2)|~r1_xboole_0(X3,esk8_3(esk1_0,X1,esk2_0))|~r1_waybel_0(esk1_0,X2,X3)|~l1_struct_0(esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_29])).
cnf(c_0_43, plain, (r1_xboole_0(X1,esk8_3(X2,X1,X3))|r2_hidden(X3,k2_pre_topc(X2,X1))|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_44, plain, ![X18]:(~l1_pre_topc(X18)|l1_struct_0(X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_45, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r1_waybel_0(X1,X2,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_46, negated_conjecture, (esk4_0=k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_47, negated_conjecture, (r2_hidden(esk2_0,k2_pre_topc(esk1_0,X1))|v2_struct_0(X2)|k10_yellow_6(esk1_0,esk3_0)!=k10_yellow_6(esk1_0,X2)|~r1_waybel_0(esk1_0,X2,X1)|~l1_struct_0(esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_38]), c_0_26]), c_0_27])]), c_0_29])).
cnf(c_0_48, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_49, negated_conjecture, (r1_waybel_0(esk1_0,esk3_0,esk4_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_23])]), c_0_28]), c_0_29])).
cnf(c_0_50, negated_conjecture, (r2_hidden(esk2_0,k2_pre_topc(esk1_0,X1))|v2_struct_0(X2)|k10_yellow_6(esk1_0,esk3_0)!=k10_yellow_6(esk1_0,X2)|~r1_waybel_0(esk1_0,X2,X1)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_26])])).
cnf(c_0_51, negated_conjecture, (r1_waybel_0(esk1_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_48]), c_0_26])])).
cnf(c_0_52, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_53, negated_conjecture, (~r2_hidden(esk2_0,k2_pre_topc(esk1_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_54, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_23]), c_0_24]), c_0_25]), c_0_52])]), c_0_53]), c_0_28]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 55
# Proof object clause steps            : 36
# Proof object formula steps           : 19
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 41
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 86
# ...of these trivial                  : 1
# ...subsumed                          : 1
# ...remaining for further processing  : 84
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 1
# Generated clauses                    : 40
# ...of the previous two non-trivial   : 33
# Contextual simplify-reflections      : 7
# Paramodulations                      : 40
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 52
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 39
# Current number of unprocessed clauses: 5
# ...number of literals in the above   : 76
# Current number of archived formulas  : 0
# Current number of archived clauses   : 32
# Clause-clause subsumption calls (NU) : 1267
# Rec. Clause-clause subsumption calls : 55
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6258

# -------------------------------------------------
# User time                : 0.014 s
# System time              : 0.000 s
# Total time               : 0.014 s
# Maximum resident set size: 28580 pages
