# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t70_yellow_0, conjecture, ![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&(v5_yellow_0(X2,X1)&m1_yellow_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>![X6]:(m1_subset_1(X6,u1_struct_0(X1))=>((X5=X3&X6=X4)=>k12_lattice3(X2,X3,X4)=k12_lattice3(X1,X5,X6)))))))), file('data/MPTP/proofs/391200df7daa4fab8698e688ece6f621.E_input', t70_yellow_0)).
fof(cc2_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v2_lattice3(X1)=>~(v2_struct_0(X1)))), file('data/MPTP/proofs/391200df7daa4fab8698e688ece6f621.E_input', cc2_lattice3)).
fof(redefinition_k7_domain_1, axiom, ![X1, X2, X3]:((~(v1_xboole_0(X1))&(m1_subset_1(X2,X1)&m1_subset_1(X3,X1)))=>k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)), file('data/MPTP/proofs/391200df7daa4fab8698e688ece6f621.E_input', redefinition_k7_domain_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/391200df7daa4fab8698e688ece6f621.E_input', t70_enumset1)).
fof(d16_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_yellow_0(X2,X1)=>(v5_yellow_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,u1_struct_0(X2))&(r2_hidden(X4,u1_struct_0(X2))&r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))))=>r2_hidden(k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4)),u1_struct_0(X2)))))))), file('data/MPTP/proofs/391200df7daa4fab8698e688ece6f621.E_input', d16_yellow_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/391200df7daa4fab8698e688ece6f621.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/391200df7daa4fab8698e688ece6f621.E_input', dt_l1_orders_2)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/391200df7daa4fab8698e688ece6f621.E_input', d2_subset_1)).
fof(dt_m1_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_yellow_0(X2,X1)=>l1_orders_2(X2))), file('data/MPTP/proofs/391200df7daa4fab8698e688ece6f621.E_input', dt_m1_yellow_0)).
fof(t21_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>(v2_lattice3(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>r2_yellow_0(X1,k2_tarski(X2,X3)))))), file('data/MPTP/proofs/391200df7daa4fab8698e688ece6f621.E_input', t21_yellow_0)).
fof(dt_k7_domain_1, axiom, ![X1, X2, X3]:((~(v1_xboole_0(X1))&(m1_subset_1(X2,X1)&m1_subset_1(X3,X1)))=>m1_subset_1(k7_domain_1(X1,X2,X3),k1_zfmisc_1(X1))), file('data/MPTP/proofs/391200df7daa4fab8698e688ece6f621.E_input', dt_k7_domain_1)).
fof(t64_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>((r2_yellow_0(X1,X3)&r2_hidden(k2_yellow_0(X1,X3),u1_struct_0(X2)))=>(r2_yellow_0(X2,X3)&k2_yellow_0(X2,X3)=k2_yellow_0(X1,X3)))))), file('data/MPTP/proofs/391200df7daa4fab8698e688ece6f621.E_input', t64_yellow_0)).
fof(cc11_yellow_0, axiom, ![X1]:((v4_orders_2(X1)&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_yellow_0(X2,X1)=>((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&v5_yellow_0(X2,X1)))=>(~(v2_struct_0(X2))&(v2_lattice3(X2)&(v4_yellow_0(X2,X1)&v5_yellow_0(X2,X1))))))), file('data/MPTP/proofs/391200df7daa4fab8698e688ece6f621.E_input', cc11_yellow_0)).
fof(cc8_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_yellow_0(X2,X1)=>(v4_yellow_0(X2,X1)=>(v5_orders_2(X2)&v4_yellow_0(X2,X1))))), file('data/MPTP/proofs/391200df7daa4fab8698e688ece6f621.E_input', cc8_yellow_0)).
fof(cc7_yellow_0, axiom, ![X1]:((v4_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_yellow_0(X2,X1)=>(v4_yellow_0(X2,X1)=>(v4_orders_2(X2)&v4_yellow_0(X2,X1))))), file('data/MPTP/proofs/391200df7daa4fab8698e688ece6f621.E_input', cc7_yellow_0)).
fof(cc6_yellow_0, axiom, ![X1]:((v3_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_yellow_0(X2,X1)=>(v4_yellow_0(X2,X1)=>(v3_orders_2(X2)&v4_yellow_0(X2,X1))))), file('data/MPTP/proofs/391200df7daa4fab8698e688ece6f621.E_input', cc6_yellow_0)).
fof(t40_yellow_0, axiom, ![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k12_lattice3(X1,X2,X3)))), file('data/MPTP/proofs/391200df7daa4fab8698e688ece6f621.E_input', t40_yellow_0)).
fof(c_0_17, negated_conjecture, ~(![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&(v5_yellow_0(X2,X1)&m1_yellow_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>![X6]:(m1_subset_1(X6,u1_struct_0(X1))=>((X5=X3&X6=X4)=>k12_lattice3(X2,X3,X4)=k12_lattice3(X1,X5,X6))))))))), inference(assume_negation,[status(cth)],[t70_yellow_0])).
fof(c_0_18, plain, ![X13]:(~l1_orders_2(X13)|(~v2_lattice3(X13)|~v2_struct_0(X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_lattice3])])])).
fof(c_0_19, negated_conjecture, ((v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v2_lattice3(esk1_0)&l1_orders_2(esk1_0)))))&((~v2_struct_0(esk2_0)&(v4_yellow_0(esk2_0,esk1_0)&(v5_yellow_0(esk2_0,esk1_0)&m1_yellow_0(esk2_0,esk1_0))))&(m1_subset_1(esk3_0,u1_struct_0(esk2_0))&(m1_subset_1(esk4_0,u1_struct_0(esk2_0))&(m1_subset_1(esk5_0,u1_struct_0(esk1_0))&(m1_subset_1(esk6_0,u1_struct_0(esk1_0))&((esk5_0=esk3_0&esk6_0=esk4_0)&k12_lattice3(esk2_0,esk3_0,esk4_0)!=k12_lattice3(esk1_0,esk5_0,esk6_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])).
fof(c_0_20, plain, ![X14, X15, X16]:(v1_xboole_0(X14)|(~m1_subset_1(X15,X14)|~m1_subset_1(X16,X14))|k7_domain_1(X14,X15,X16)=k2_tarski(X15,X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k7_domain_1])])])).
fof(c_0_21, plain, ![X17, X18]:k1_enumset1(X17,X17,X18)=k2_tarski(X17,X18), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_22, plain, (~l1_orders_2(X1)|~v2_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, negated_conjecture, (v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, plain, (v1_xboole_0(X1)|k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_26, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, negated_conjecture, (esk6_0=esk4_0), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_28, plain, ![X19, X20, X21, X22]:((~v5_yellow_0(X20,X19)|(~m1_subset_1(X21,u1_struct_0(X19))|(~m1_subset_1(X22,u1_struct_0(X19))|(~r2_hidden(X21,u1_struct_0(X20))|(~r2_hidden(X22,u1_struct_0(X20))|~r2_yellow_0(X19,k7_domain_1(u1_struct_0(X19),X21,X22)))|r2_hidden(k2_yellow_0(X19,k7_domain_1(u1_struct_0(X19),X21,X22)),u1_struct_0(X20)))))|~m1_yellow_0(X20,X19)|(v2_struct_0(X19)|~l1_orders_2(X19)))&((m1_subset_1(esk7_2(X19,X20),u1_struct_0(X19))|v5_yellow_0(X20,X19)|~m1_yellow_0(X20,X19)|(v2_struct_0(X19)|~l1_orders_2(X19)))&((m1_subset_1(esk8_2(X19,X20),u1_struct_0(X19))|v5_yellow_0(X20,X19)|~m1_yellow_0(X20,X19)|(v2_struct_0(X19)|~l1_orders_2(X19)))&(((r2_hidden(esk7_2(X19,X20),u1_struct_0(X20))|v5_yellow_0(X20,X19)|~m1_yellow_0(X20,X19)|(v2_struct_0(X19)|~l1_orders_2(X19)))&((r2_hidden(esk8_2(X19,X20),u1_struct_0(X20))|v5_yellow_0(X20,X19)|~m1_yellow_0(X20,X19)|(v2_struct_0(X19)|~l1_orders_2(X19)))&(r2_yellow_0(X19,k7_domain_1(u1_struct_0(X19),esk7_2(X19,X20),esk8_2(X19,X20)))|v5_yellow_0(X20,X19)|~m1_yellow_0(X20,X19)|(v2_struct_0(X19)|~l1_orders_2(X19)))))&(~r2_hidden(k2_yellow_0(X19,k7_domain_1(u1_struct_0(X19),esk7_2(X19,X20),esk8_2(X19,X20))),u1_struct_0(X20))|v5_yellow_0(X20,X19)|~m1_yellow_0(X20,X19)|(v2_struct_0(X19)|~l1_orders_2(X19))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_yellow_0])])])])])])).
cnf(c_0_29, negated_conjecture, (~v2_struct_0(esk1_0)|~l1_orders_2(esk1_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_30, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_31, plain, ![X25]:(v2_struct_0(X25)|~l1_struct_0(X25)|~v1_xboole_0(u1_struct_0(X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_32, plain, ![X26]:(~l1_orders_2(X26)|l1_struct_0(X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
fof(c_0_33, plain, ![X27, X28]:(((~m1_subset_1(X28,X27)|r2_hidden(X28,X27)|v1_xboole_0(X27))&(~r2_hidden(X28,X27)|m1_subset_1(X28,X27)|v1_xboole_0(X27)))&((~m1_subset_1(X28,X27)|v1_xboole_0(X28)|~v1_xboole_0(X27))&(~v1_xboole_0(X28)|m1_subset_1(X28,X27)|~v1_xboole_0(X27)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_34, plain, ![X29, X30]:(~l1_orders_2(X29)|(~m1_yellow_0(X30,X29)|l1_orders_2(X30))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_yellow_0])])])).
fof(c_0_35, plain, ![X31, X32, X33]:((~v2_lattice3(X31)|(~m1_subset_1(X32,u1_struct_0(X31))|(~m1_subset_1(X33,u1_struct_0(X31))|r2_yellow_0(X31,k2_tarski(X32,X33))))|(~v5_orders_2(X31)|~l1_orders_2(X31)))&((m1_subset_1(esk9_1(X31),u1_struct_0(X31))|v2_lattice3(X31)|(~v5_orders_2(X31)|~l1_orders_2(X31)))&((m1_subset_1(esk10_1(X31),u1_struct_0(X31))|v2_lattice3(X31)|(~v5_orders_2(X31)|~l1_orders_2(X31)))&(~r2_yellow_0(X31,k2_tarski(esk9_1(X31),esk10_1(X31)))|v2_lattice3(X31)|(~v5_orders_2(X31)|~l1_orders_2(X31)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t21_yellow_0])])])])])).
cnf(c_0_36, plain, (k1_enumset1(X2,X2,X3)=k7_domain_1(X1,X2,X3)|v1_xboole_0(X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_37, negated_conjecture, (m1_subset_1(esk6_0,u1_struct_0(esk2_0))), inference(rw,[status(thm)],[c_0_26, c_0_27])).
fof(c_0_38, plain, ![X36, X37, X38]:(v1_xboole_0(X36)|(~m1_subset_1(X37,X36)|~m1_subset_1(X38,X36))|m1_subset_1(k7_domain_1(X36,X37,X38),k1_zfmisc_1(X36))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_domain_1])])])).
cnf(c_0_39, plain, (r2_hidden(k2_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4)),u1_struct_0(X1))|v2_struct_0(X2)|~v5_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,u1_struct_0(X1))|~r2_hidden(X4,u1_struct_0(X1))|~r2_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))|~m1_yellow_0(X1,X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_40, negated_conjecture, (m1_subset_1(esk6_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_41, negated_conjecture, (~v2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_30])])).
cnf(c_0_42, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_43, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_44, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_45, plain, (l1_orders_2(X2)|~l1_orders_2(X1)|~m1_yellow_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_46, negated_conjecture, (m1_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_47, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|~v2_lattice3(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_48, negated_conjecture, (k1_enumset1(X1,X1,esk6_0)=k7_domain_1(u1_struct_0(esk2_0),X1,esk6_0)|v1_xboole_0(u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_49, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_50, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_51, negated_conjecture, (esk5_0=esk3_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_52, plain, (v1_xboole_0(X1)|m1_subset_1(k7_domain_1(X1,X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_53, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_54, negated_conjecture, (r2_hidden(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,esk6_0)),u1_struct_0(X2))|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,esk6_0))|~r2_hidden(esk6_0,u1_struct_0(X2))|~r2_hidden(X1,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_yellow_0(X2,esk1_0)|~v5_yellow_0(X2,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_30])]), c_0_41])).
cnf(c_0_55, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_56, negated_conjecture, (r2_hidden(esk6_0,u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_44, c_0_37])).
cnf(c_0_57, negated_conjecture, (l1_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_30])])).
cnf(c_0_58, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_59, plain, (r2_yellow_0(X1,k1_enumset1(X2,X2,X3))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[c_0_47, c_0_25])).
cnf(c_0_60, negated_conjecture, (k1_enumset1(esk3_0,esk3_0,esk6_0)=k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0)|v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_61, negated_conjecture, (k1_enumset1(X1,X1,esk6_0)=k7_domain_1(u1_struct_0(esk1_0),X1,esk6_0)|v1_xboole_0(u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_36, c_0_40])).
cnf(c_0_62, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(rw,[status(thm)],[c_0_50, c_0_51])).
fof(c_0_63, plain, ![X39, X40, X41]:((r2_yellow_0(X40,X41)|(~r2_yellow_0(X39,X41)|~r2_hidden(k2_yellow_0(X39,X41),u1_struct_0(X40)))|~m1_subset_1(X41,k1_zfmisc_1(u1_struct_0(X40)))|(v2_struct_0(X40)|(~v4_yellow_0(X40,X39)|~m1_yellow_0(X40,X39)))|(v2_struct_0(X39)|(~v4_orders_2(X39)|~l1_orders_2(X39))))&(k2_yellow_0(X40,X41)=k2_yellow_0(X39,X41)|(~r2_yellow_0(X39,X41)|~r2_hidden(k2_yellow_0(X39,X41),u1_struct_0(X40)))|~m1_subset_1(X41,k1_zfmisc_1(u1_struct_0(X40)))|(v2_struct_0(X40)|(~v4_yellow_0(X40,X39)|~m1_yellow_0(X40,X39)))|(v2_struct_0(X39)|(~v4_orders_2(X39)|~l1_orders_2(X39))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t64_yellow_0])])])])])).
cnf(c_0_64, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|m1_subset_1(k7_domain_1(u1_struct_0(esk2_0),X1,esk6_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_52, c_0_37])).
cnf(c_0_65, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,esk6_0)),u1_struct_0(X1))|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,esk6_0))|~r2_hidden(esk6_0,u1_struct_0(X1))|~r2_hidden(X2,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_yellow_0(X1,esk1_0)|~v5_yellow_0(X1,esk1_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_66, negated_conjecture, (r2_hidden(esk6_0,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57])]), c_0_58])).
cnf(c_0_67, negated_conjecture, (v5_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_68, negated_conjecture, (r2_hidden(esk3_0,u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_44, c_0_49])).
cnf(c_0_69, negated_conjecture, (r2_yellow_0(X1,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))|v1_xboole_0(u1_struct_0(esk2_0))|~m1_subset_1(esk6_0,u1_struct_0(X1))|~m1_subset_1(esk3_0,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_70, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_71, negated_conjecture, (k1_enumset1(esk3_0,esk3_0,esk6_0)=k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_72, plain, (k2_yellow_0(X1,X2)=k2_yellow_0(X3,X2)|v2_struct_0(X1)|v2_struct_0(X3)|~r2_yellow_0(X3,X2)|~r2_hidden(k2_yellow_0(X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v4_yellow_0(X1,X3)|~m1_yellow_0(X1,X3)|~v4_orders_2(X3)|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_73, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|m1_subset_1(k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_64, c_0_49])).
cnf(c_0_74, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|m1_subset_1(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,esk6_0)),u1_struct_0(esk2_0))|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,esk6_0))|~r2_hidden(X1,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_46]), c_0_66]), c_0_67])])).
cnf(c_0_75, negated_conjecture, (r2_hidden(esk3_0,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_68]), c_0_57])]), c_0_58])).
cnf(c_0_76, negated_conjecture, (r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_40]), c_0_62]), c_0_30]), c_0_23]), c_0_70])])).
cnf(c_0_77, negated_conjecture, (k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0)=k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)|v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_60, c_0_71])).
fof(c_0_78, plain, ![X42, X43]:((~v2_struct_0(X43)|(v2_struct_0(X43)|(~v4_yellow_0(X43,X42)|~v5_yellow_0(X43,X42)))|~m1_yellow_0(X43,X42)|(~v4_orders_2(X42)|(~v5_orders_2(X42)|(~v2_lattice3(X42)|~l1_orders_2(X42)))))&((v2_lattice3(X43)|(v2_struct_0(X43)|(~v4_yellow_0(X43,X42)|~v5_yellow_0(X43,X42)))|~m1_yellow_0(X43,X42)|(~v4_orders_2(X42)|(~v5_orders_2(X42)|(~v2_lattice3(X42)|~l1_orders_2(X42)))))&((v4_yellow_0(X43,X42)|(v2_struct_0(X43)|(~v4_yellow_0(X43,X42)|~v5_yellow_0(X43,X42)))|~m1_yellow_0(X43,X42)|(~v4_orders_2(X42)|(~v5_orders_2(X42)|(~v2_lattice3(X42)|~l1_orders_2(X42)))))&(v5_yellow_0(X43,X42)|(v2_struct_0(X43)|(~v4_yellow_0(X43,X42)|~v5_yellow_0(X43,X42)))|~m1_yellow_0(X43,X42)|(~v4_orders_2(X42)|(~v5_orders_2(X42)|(~v2_lattice3(X42)|~l1_orders_2(X42)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc11_yellow_0])])])])])).
fof(c_0_79, plain, ![X44, X45]:((v5_orders_2(X45)|~v4_yellow_0(X45,X44)|~m1_yellow_0(X45,X44)|(~v5_orders_2(X44)|~l1_orders_2(X44)))&(v4_yellow_0(X45,X44)|~v4_yellow_0(X45,X44)|~m1_yellow_0(X45,X44)|(~v5_orders_2(X44)|~l1_orders_2(X44)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_yellow_0])])])])).
fof(c_0_80, plain, ![X46, X47]:((v4_orders_2(X47)|~v4_yellow_0(X47,X46)|~m1_yellow_0(X47,X46)|(~v4_orders_2(X46)|~l1_orders_2(X46)))&(v4_yellow_0(X47,X46)|~v4_yellow_0(X47,X46)|~m1_yellow_0(X47,X46)|(~v4_orders_2(X46)|~l1_orders_2(X46)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc7_yellow_0])])])])).
fof(c_0_81, plain, ![X48, X49]:((v3_orders_2(X49)|~v4_yellow_0(X49,X48)|~m1_yellow_0(X49,X48)|(~v3_orders_2(X48)|~l1_orders_2(X48)))&(v4_yellow_0(X49,X48)|~v4_yellow_0(X49,X48)|~m1_yellow_0(X49,X48)|(~v3_orders_2(X48)|~l1_orders_2(X48)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc6_yellow_0])])])])).
cnf(c_0_82, negated_conjecture, (k2_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))=k2_yellow_0(X1,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))|v1_xboole_0(u1_struct_0(esk2_0))|v2_struct_0(X1)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))|~r2_hidden(k2_yellow_0(X1,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0)),u1_struct_0(esk2_0))|~m1_yellow_0(esk2_0,X1)|~v4_yellow_0(esk2_0,X1)|~l1_orders_2(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_58])).
cnf(c_0_83, negated_conjecture, (v4_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_84, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_85, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|m1_subset_1(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)),u1_struct_0(esk2_0))|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_62])])).
cnf(c_0_86, negated_conjecture, (r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))|v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_76, c_0_77])).
fof(c_0_87, plain, ![X50, X51, X52]:(~v3_orders_2(X50)|(~v4_orders_2(X50)|(~v5_orders_2(X50)|(~v2_lattice3(X50)|~l1_orders_2(X50))))|(~m1_subset_1(X51,u1_struct_0(X50))|(~m1_subset_1(X52,u1_struct_0(X50))|k2_yellow_0(X50,k7_domain_1(u1_struct_0(X50),X51,X52))=k12_lattice3(X50,X51,X52)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t40_yellow_0])])])).
cnf(c_0_88, plain, (v2_lattice3(X1)|v2_struct_0(X1)|~v4_yellow_0(X1,X2)|~v5_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v2_lattice3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_89, plain, (v5_orders_2(X1)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_90, plain, (v4_orders_2(X1)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~v4_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_91, plain, (v3_orders_2(X1)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~v3_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_92, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_93, negated_conjecture, (k2_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))=k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))|v1_xboole_0(u1_struct_0(esk2_0))|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))|~r2_hidden(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0)),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_46]), c_0_83]), c_0_30]), c_0_84])]), c_0_41])).
cnf(c_0_94, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))|m1_subset_1(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_95, plain, (k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k12_lattice3(X1,X2,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_96, negated_conjecture, (v2_lattice3(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_46]), c_0_67]), c_0_83]), c_0_30]), c_0_23]), c_0_70]), c_0_84])]), c_0_58])).
cnf(c_0_97, negated_conjecture, (v5_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_83]), c_0_46]), c_0_30]), c_0_70])])).
cnf(c_0_98, negated_conjecture, (v4_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_83]), c_0_46]), c_0_30]), c_0_84])])).
cnf(c_0_99, negated_conjecture, (v3_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_46]), c_0_83]), c_0_30]), c_0_92])])).
cnf(c_0_100, negated_conjecture, (k2_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))=k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))|v1_xboole_0(u1_struct_0(esk2_0))|~r2_hidden(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0)),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_93, c_0_76])).
cnf(c_0_101, negated_conjecture, (r2_hidden(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)),u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_44, c_0_94])).
cnf(c_0_102, negated_conjecture, (k2_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))=k12_lattice3(esk2_0,esk3_0,esk6_0)|v1_xboole_0(u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_77]), c_0_37]), c_0_49]), c_0_57]), c_0_96]), c_0_97]), c_0_98]), c_0_99])])).
cnf(c_0_103, negated_conjecture, (k2_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))=k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))|v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_77]), c_0_101])).
cnf(c_0_104, negated_conjecture, (k12_lattice3(esk2_0,esk3_0,esk4_0)!=k12_lattice3(esk1_0,esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_105, negated_conjecture, (k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))=k12_lattice3(esk2_0,esk3_0,esk6_0)|v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_106, negated_conjecture, (k12_lattice3(esk2_0,esk3_0,esk6_0)!=k12_lattice3(esk1_0,esk3_0,esk6_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_104, c_0_27]), c_0_51])).
cnf(c_0_107, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_105]), c_0_40]), c_0_62]), c_0_30]), c_0_23]), c_0_70]), c_0_84]), c_0_92])]), c_0_106])).
cnf(c_0_108, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_107]), c_0_57])]), c_0_58])).
cnf(c_0_109, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_108]), c_0_30])]), c_0_41]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 110
# Proof object clause steps            : 75
# Proof object formula steps           : 35
# Proof object conjectures             : 58
# Proof object clause conjectures      : 55
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 33
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 36
# Proof object simplifying inferences  : 79
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 51
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 44
# Processed clauses                    : 651
# ...of these trivial                  : 0
# ...subsumed                          : 125
# ...remaining for further processing  : 526
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 284
# Backward-rewritten                   : 50
# Generated clauses                    : 1703
# ...of the previous two non-trivial   : 1487
# Contextual simplify-reflections      : 35
# Paramodulations                      : 1703
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 148
#    Positive orientable unit clauses  : 24
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 121
# Current number of unprocessed clauses: 911
# ...number of literals in the above   : 4475
# Current number of archived formulas  : 0
# Current number of archived clauses   : 379
# Clause-clause subsumption calls (NU) : 41528
# Rec. Clause-clause subsumption calls : 7648
# Non-unit clause-clause subsumptions  : 435
# Unit Clause-clause subsumption calls : 55
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 71589

# -------------------------------------------------
# User time                : 0.051 s
# System time              : 0.007 s
# Total time               : 0.058 s
# Maximum resident set size: 29020 pages
