# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t70_yellow_0, conjecture, ![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&(v5_yellow_0(X2,X1)&m1_yellow_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>![X6]:(m1_subset_1(X6,u1_struct_0(X1))=>((X5=X3&X6=X4)=>k12_lattice3(X2,X3,X4)=k12_lattice3(X1,X5,X6)))))))), file('data/MPTP/proofs/39e00780a831490c84bace181131c633.E_input', t70_yellow_0)).
fof(t21_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>(v2_lattice3(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>r2_yellow_0(X1,k2_tarski(X2,X3)))))), file('data/MPTP/proofs/39e00780a831490c84bace181131c633.E_input', t21_yellow_0)).
fof(t77_enumset1, axiom, ![X1, X2]:k2_enumset1(X1,X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/39e00780a831490c84bace181131c633.E_input', t77_enumset1)).
fof(redefinition_k7_domain_1, axiom, ![X1, X2, X3]:((~(v1_xboole_0(X1))&(m1_subset_1(X2,X1)&m1_subset_1(X3,X1)))=>k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)), file('data/MPTP/proofs/39e00780a831490c84bace181131c633.E_input', redefinition_k7_domain_1)).
fof(dt_m1_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_yellow_0(X2,X1)=>l1_orders_2(X2))), file('data/MPTP/proofs/39e00780a831490c84bace181131c633.E_input', dt_m1_yellow_0)).
fof(cc11_yellow_0, axiom, ![X1]:((v4_orders_2(X1)&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_yellow_0(X2,X1)=>((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&v5_yellow_0(X2,X1)))=>(~(v2_struct_0(X2))&(v2_lattice3(X2)&(v4_yellow_0(X2,X1)&v5_yellow_0(X2,X1))))))), file('data/MPTP/proofs/39e00780a831490c84bace181131c633.E_input', cc11_yellow_0)).
fof(cc8_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_yellow_0(X2,X1)=>(v4_yellow_0(X2,X1)=>(v5_orders_2(X2)&v4_yellow_0(X2,X1))))), file('data/MPTP/proofs/39e00780a831490c84bace181131c633.E_input', cc8_yellow_0)).
fof(cc7_yellow_0, axiom, ![X1]:((v4_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_yellow_0(X2,X1)=>(v4_yellow_0(X2,X1)=>(v4_orders_2(X2)&v4_yellow_0(X2,X1))))), file('data/MPTP/proofs/39e00780a831490c84bace181131c633.E_input', cc7_yellow_0)).
fof(cc6_yellow_0, axiom, ![X1]:((v3_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_yellow_0(X2,X1)=>(v4_yellow_0(X2,X1)=>(v3_orders_2(X2)&v4_yellow_0(X2,X1))))), file('data/MPTP/proofs/39e00780a831490c84bace181131c633.E_input', cc6_yellow_0)).
fof(t40_yellow_0, axiom, ![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k12_lattice3(X1,X2,X3)))), file('data/MPTP/proofs/39e00780a831490c84bace181131c633.E_input', t40_yellow_0)).
fof(cc2_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v2_lattice3(X1)=>~(v2_struct_0(X1)))), file('data/MPTP/proofs/39e00780a831490c84bace181131c633.E_input', cc2_lattice3)).
fof(t64_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>((r2_yellow_0(X1,X3)&r2_hidden(k2_yellow_0(X1,X3),u1_struct_0(X2)))=>(r2_yellow_0(X2,X3)&k2_yellow_0(X2,X3)=k2_yellow_0(X1,X3)))))), file('data/MPTP/proofs/39e00780a831490c84bace181131c633.E_input', t64_yellow_0)).
fof(dt_k7_domain_1, axiom, ![X1, X2, X3]:((~(v1_xboole_0(X1))&(m1_subset_1(X2,X1)&m1_subset_1(X3,X1)))=>m1_subset_1(k7_domain_1(X1,X2,X3),k1_zfmisc_1(X1))), file('data/MPTP/proofs/39e00780a831490c84bace181131c633.E_input', dt_k7_domain_1)).
fof(d16_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_yellow_0(X2,X1)=>(v5_yellow_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,u1_struct_0(X2))&(r2_hidden(X4,u1_struct_0(X2))&r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))))=>r2_hidden(k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4)),u1_struct_0(X2)))))))), file('data/MPTP/proofs/39e00780a831490c84bace181131c633.E_input', d16_yellow_0)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/39e00780a831490c84bace181131c633.E_input', t2_subset)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/39e00780a831490c84bace181131c633.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/39e00780a831490c84bace181131c633.E_input', dt_l1_orders_2)).
fof(c_0_17, negated_conjecture, ~(![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&(v5_yellow_0(X2,X1)&m1_yellow_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>![X6]:(m1_subset_1(X6,u1_struct_0(X1))=>((X5=X3&X6=X4)=>k12_lattice3(X2,X3,X4)=k12_lattice3(X1,X5,X6))))))))), inference(assume_negation,[status(cth)],[t70_yellow_0])).
fof(c_0_18, plain, ![X34, X35, X36]:((~v2_lattice3(X34)|(~m1_subset_1(X35,u1_struct_0(X34))|(~m1_subset_1(X36,u1_struct_0(X34))|r2_yellow_0(X34,k2_tarski(X35,X36))))|(~v5_orders_2(X34)|~l1_orders_2(X34)))&((m1_subset_1(esk9_1(X34),u1_struct_0(X34))|v2_lattice3(X34)|(~v5_orders_2(X34)|~l1_orders_2(X34)))&((m1_subset_1(esk10_1(X34),u1_struct_0(X34))|v2_lattice3(X34)|(~v5_orders_2(X34)|~l1_orders_2(X34)))&(~r2_yellow_0(X34,k2_tarski(esk9_1(X34),esk10_1(X34)))|v2_lattice3(X34)|(~v5_orders_2(X34)|~l1_orders_2(X34)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t21_yellow_0])])])])])).
fof(c_0_19, plain, ![X14, X15]:k2_enumset1(X14,X14,X14,X15)=k2_tarski(X14,X15), inference(variable_rename,[status(thm)],[t77_enumset1])).
fof(c_0_20, plain, ![X44, X45, X46]:(v1_xboole_0(X44)|(~m1_subset_1(X45,X44)|~m1_subset_1(X46,X44))|k7_domain_1(X44,X45,X46)=k2_tarski(X45,X46)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k7_domain_1])])])).
fof(c_0_21, negated_conjecture, ((v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v2_lattice3(esk1_0)&l1_orders_2(esk1_0)))))&((~v2_struct_0(esk2_0)&(v4_yellow_0(esk2_0,esk1_0)&(v5_yellow_0(esk2_0,esk1_0)&m1_yellow_0(esk2_0,esk1_0))))&(m1_subset_1(esk3_0,u1_struct_0(esk2_0))&(m1_subset_1(esk4_0,u1_struct_0(esk2_0))&(m1_subset_1(esk5_0,u1_struct_0(esk1_0))&(m1_subset_1(esk6_0,u1_struct_0(esk1_0))&((esk5_0=esk3_0&esk6_0=esk4_0)&k12_lattice3(esk2_0,esk3_0,esk4_0)!=k12_lattice3(esk1_0,esk5_0,esk6_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])).
fof(c_0_22, plain, ![X25, X26]:(~l1_orders_2(X25)|(~m1_yellow_0(X26,X25)|l1_orders_2(X26))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_yellow_0])])])).
fof(c_0_23, plain, ![X49, X50]:((~v2_struct_0(X50)|(v2_struct_0(X50)|(~v4_yellow_0(X50,X49)|~v5_yellow_0(X50,X49)))|~m1_yellow_0(X50,X49)|(~v4_orders_2(X49)|(~v5_orders_2(X49)|(~v2_lattice3(X49)|~l1_orders_2(X49)))))&((v2_lattice3(X50)|(v2_struct_0(X50)|(~v4_yellow_0(X50,X49)|~v5_yellow_0(X50,X49)))|~m1_yellow_0(X50,X49)|(~v4_orders_2(X49)|(~v5_orders_2(X49)|(~v2_lattice3(X49)|~l1_orders_2(X49)))))&((v4_yellow_0(X50,X49)|(v2_struct_0(X50)|(~v4_yellow_0(X50,X49)|~v5_yellow_0(X50,X49)))|~m1_yellow_0(X50,X49)|(~v4_orders_2(X49)|(~v5_orders_2(X49)|(~v2_lattice3(X49)|~l1_orders_2(X49)))))&(v5_yellow_0(X50,X49)|(v2_struct_0(X50)|(~v4_yellow_0(X50,X49)|~v5_yellow_0(X50,X49)))|~m1_yellow_0(X50,X49)|(~v4_orders_2(X49)|(~v5_orders_2(X49)|(~v2_lattice3(X49)|~l1_orders_2(X49)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc11_yellow_0])])])])])).
fof(c_0_24, plain, ![X22, X23]:((v5_orders_2(X23)|~v4_yellow_0(X23,X22)|~m1_yellow_0(X23,X22)|(~v5_orders_2(X22)|~l1_orders_2(X22)))&(v4_yellow_0(X23,X22)|~v4_yellow_0(X23,X22)|~m1_yellow_0(X23,X22)|(~v5_orders_2(X22)|~l1_orders_2(X22)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_yellow_0])])])])).
fof(c_0_25, plain, ![X39, X40]:((v4_orders_2(X40)|~v4_yellow_0(X40,X39)|~m1_yellow_0(X40,X39)|(~v4_orders_2(X39)|~l1_orders_2(X39)))&(v4_yellow_0(X40,X39)|~v4_yellow_0(X40,X39)|~m1_yellow_0(X40,X39)|(~v4_orders_2(X39)|~l1_orders_2(X39)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc7_yellow_0])])])])).
fof(c_0_26, plain, ![X47, X48]:((v3_orders_2(X48)|~v4_yellow_0(X48,X47)|~m1_yellow_0(X48,X47)|(~v3_orders_2(X47)|~l1_orders_2(X47)))&(v4_yellow_0(X48,X47)|~v4_yellow_0(X48,X47)|~m1_yellow_0(X48,X47)|(~v3_orders_2(X47)|~l1_orders_2(X47)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc6_yellow_0])])])])).
cnf(c_0_27, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|~v2_lattice3(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, plain, (k2_enumset1(X1,X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (v1_xboole_0(X1)|k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_30, plain, ![X31, X32, X33]:(~v3_orders_2(X31)|(~v4_orders_2(X31)|(~v5_orders_2(X31)|(~v2_lattice3(X31)|~l1_orders_2(X31))))|(~m1_subset_1(X32,u1_struct_0(X31))|(~m1_subset_1(X33,u1_struct_0(X31))|k2_yellow_0(X31,k7_domain_1(u1_struct_0(X31),X32,X33))=k12_lattice3(X31,X32,X33)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t40_yellow_0])])])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_32, negated_conjecture, (esk6_0=esk4_0), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_33, plain, (l1_orders_2(X2)|~l1_orders_2(X1)|~m1_yellow_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_34, negated_conjecture, (m1_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_35, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_36, plain, (v2_lattice3(X1)|v2_struct_0(X1)|~v4_yellow_0(X1,X2)|~v5_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v2_lattice3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_37, negated_conjecture, (v4_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_38, negated_conjecture, (v5_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_39, negated_conjecture, (v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_40, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_41, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_42, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_43, plain, (v5_orders_2(X1)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_44, plain, (v4_orders_2(X1)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~v4_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_45, plain, (v3_orders_2(X1)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~v3_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_46, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_47, plain, (r2_yellow_0(X1,k2_enumset1(X2,X2,X2,X3))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_48, negated_conjecture, (m1_subset_1(esk6_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_49, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_50, negated_conjecture, (esk5_0=esk3_0), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_51, plain, ![X27]:(~l1_orders_2(X27)|(~v2_lattice3(X27)|~v2_struct_0(X27))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_lattice3])])])).
cnf(c_0_52, plain, (k7_domain_1(X1,X2,X3)=k2_enumset1(X2,X2,X2,X3)|v1_xboole_0(X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_29, c_0_28])).
cnf(c_0_53, plain, (k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k12_lattice3(X1,X2,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(esk6_0,u1_struct_0(esk2_0))), inference(rw,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_55, negated_conjecture, (l1_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])])).
cnf(c_0_56, negated_conjecture, (v2_lattice3(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_34]), c_0_38]), c_0_35]), c_0_39]), c_0_40]), c_0_41])]), c_0_42])).
cnf(c_0_57, negated_conjecture, (v5_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_37]), c_0_34]), c_0_35]), c_0_40])])).
cnf(c_0_58, negated_conjecture, (v4_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_37]), c_0_34]), c_0_35]), c_0_41])])).
cnf(c_0_59, negated_conjecture, (v3_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_37]), c_0_34]), c_0_35]), c_0_46])])).
fof(c_0_60, plain, ![X28, X29, X30]:((r2_yellow_0(X29,X30)|(~r2_yellow_0(X28,X30)|~r2_hidden(k2_yellow_0(X28,X30),u1_struct_0(X29)))|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))|(v2_struct_0(X29)|(~v4_yellow_0(X29,X28)|~m1_yellow_0(X29,X28)))|(v2_struct_0(X28)|(~v4_orders_2(X28)|~l1_orders_2(X28))))&(k2_yellow_0(X29,X30)=k2_yellow_0(X28,X30)|(~r2_yellow_0(X28,X30)|~r2_hidden(k2_yellow_0(X28,X30),u1_struct_0(X29)))|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))|(v2_struct_0(X29)|(~v4_yellow_0(X29,X28)|~m1_yellow_0(X29,X28)))|(v2_struct_0(X28)|(~v4_orders_2(X28)|~l1_orders_2(X28))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t64_yellow_0])])])])])).
cnf(c_0_61, negated_conjecture, (r2_yellow_0(esk1_0,k2_enumset1(X1,X1,X1,esk6_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_35]), c_0_39]), c_0_40])])).
cnf(c_0_62, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(rw,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_63, plain, (~l1_orders_2(X1)|~v2_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
fof(c_0_64, plain, ![X41, X42, X43]:(v1_xboole_0(X41)|(~m1_subset_1(X42,X41)|~m1_subset_1(X43,X41))|m1_subset_1(k7_domain_1(X41,X42,X43),k1_zfmisc_1(X41))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_domain_1])])])).
cnf(c_0_65, negated_conjecture, (k2_enumset1(X1,X1,X1,esk6_0)=k7_domain_1(u1_struct_0(esk1_0),X1,esk6_0)|v1_xboole_0(u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_52, c_0_48])).
cnf(c_0_66, negated_conjecture, (k12_lattice3(esk2_0,esk3_0,esk4_0)!=k12_lattice3(esk1_0,esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_67, negated_conjecture, (k12_lattice3(esk2_0,X1,esk6_0)=k2_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),X1,esk6_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55]), c_0_56]), c_0_57]), c_0_58]), c_0_59])])).
cnf(c_0_68, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_69, plain, (k2_yellow_0(X1,X2)=k2_yellow_0(X3,X2)|v2_struct_0(X1)|v2_struct_0(X3)|~r2_yellow_0(X3,X2)|~r2_hidden(k2_yellow_0(X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v4_yellow_0(X1,X3)|~m1_yellow_0(X1,X3)|~v4_orders_2(X3)|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_70, negated_conjecture, (r2_yellow_0(esk1_0,k2_enumset1(esk3_0,esk3_0,esk3_0,esk6_0))), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_71, negated_conjecture, (~v2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_39]), c_0_35])])).
cnf(c_0_72, plain, (v1_xboole_0(X1)|m1_subset_1(k7_domain_1(X1,X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
fof(c_0_73, plain, ![X16, X17, X18, X19]:((~v5_yellow_0(X17,X16)|(~m1_subset_1(X18,u1_struct_0(X16))|(~m1_subset_1(X19,u1_struct_0(X16))|(~r2_hidden(X18,u1_struct_0(X17))|(~r2_hidden(X19,u1_struct_0(X17))|~r2_yellow_0(X16,k7_domain_1(u1_struct_0(X16),X18,X19)))|r2_hidden(k2_yellow_0(X16,k7_domain_1(u1_struct_0(X16),X18,X19)),u1_struct_0(X17)))))|~m1_yellow_0(X17,X16)|(v2_struct_0(X16)|~l1_orders_2(X16)))&((m1_subset_1(esk7_2(X16,X17),u1_struct_0(X16))|v5_yellow_0(X17,X16)|~m1_yellow_0(X17,X16)|(v2_struct_0(X16)|~l1_orders_2(X16)))&((m1_subset_1(esk8_2(X16,X17),u1_struct_0(X16))|v5_yellow_0(X17,X16)|~m1_yellow_0(X17,X16)|(v2_struct_0(X16)|~l1_orders_2(X16)))&(((r2_hidden(esk7_2(X16,X17),u1_struct_0(X17))|v5_yellow_0(X17,X16)|~m1_yellow_0(X17,X16)|(v2_struct_0(X16)|~l1_orders_2(X16)))&((r2_hidden(esk8_2(X16,X17),u1_struct_0(X17))|v5_yellow_0(X17,X16)|~m1_yellow_0(X17,X16)|(v2_struct_0(X16)|~l1_orders_2(X16)))&(r2_yellow_0(X16,k7_domain_1(u1_struct_0(X16),esk7_2(X16,X17),esk8_2(X16,X17)))|v5_yellow_0(X17,X16)|~m1_yellow_0(X17,X16)|(v2_struct_0(X16)|~l1_orders_2(X16)))))&(~r2_hidden(k2_yellow_0(X16,k7_domain_1(u1_struct_0(X16),esk7_2(X16,X17),esk8_2(X16,X17))),u1_struct_0(X17))|v5_yellow_0(X17,X16)|~m1_yellow_0(X17,X16)|(v2_struct_0(X16)|~l1_orders_2(X16))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_yellow_0])])])])])])).
cnf(c_0_74, negated_conjecture, (k2_enumset1(esk3_0,esk3_0,esk3_0,esk6_0)=k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_65, c_0_62])).
fof(c_0_75, plain, ![X51, X52]:(~m1_subset_1(X51,X52)|(v1_xboole_0(X52)|r2_hidden(X51,X52))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_76, negated_conjecture, (k12_lattice3(esk2_0,esk3_0,esk6_0)!=k12_lattice3(esk1_0,esk3_0,esk6_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_32]), c_0_50])).
cnf(c_0_77, negated_conjecture, (k12_lattice3(esk2_0,esk3_0,esk6_0)=k2_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_78, negated_conjecture, (k12_lattice3(esk1_0,X1,esk6_0)=k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,esk6_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_48]), c_0_35]), c_0_39]), c_0_40]), c_0_41]), c_0_46])])).
cnf(c_0_79, negated_conjecture, (k2_enumset1(X1,X1,X1,esk6_0)=k7_domain_1(u1_struct_0(esk2_0),X1,esk6_0)|v1_xboole_0(u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_52, c_0_54])).
cnf(c_0_80, negated_conjecture, (k2_yellow_0(X1,k2_enumset1(esk3_0,esk3_0,esk3_0,esk6_0))=k2_yellow_0(esk1_0,k2_enumset1(esk3_0,esk3_0,esk3_0,esk6_0))|v2_struct_0(X1)|~r2_hidden(k2_yellow_0(esk1_0,k2_enumset1(esk3_0,esk3_0,esk3_0,esk6_0)),u1_struct_0(X1))|~m1_subset_1(k2_enumset1(esk3_0,esk3_0,esk3_0,esk6_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_yellow_0(X1,esk1_0)|~v4_yellow_0(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_35]), c_0_41])]), c_0_71])).
cnf(c_0_81, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|m1_subset_1(k7_domain_1(u1_struct_0(esk2_0),X1,esk6_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_72, c_0_54])).
cnf(c_0_82, plain, (r2_hidden(k2_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4)),u1_struct_0(X1))|v2_struct_0(X2)|~v5_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,u1_struct_0(X1))|~r2_hidden(X4,u1_struct_0(X1))|~r2_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))|~m1_yellow_0(X1,X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_83, negated_conjecture, (r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_70, c_0_74])).
cnf(c_0_84, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_85, negated_conjecture, (k12_lattice3(esk1_0,esk3_0,esk6_0)!=k2_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))), inference(rw,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_86, negated_conjecture, (k12_lattice3(esk1_0,esk3_0,esk6_0)=k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))), inference(spm,[status(thm)],[c_0_78, c_0_62])).
cnf(c_0_87, negated_conjecture, (k2_enumset1(esk3_0,esk3_0,esk3_0,esk6_0)=k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0)|v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_79, c_0_68])).
cnf(c_0_88, negated_conjecture, (k2_yellow_0(esk2_0,k2_enumset1(esk3_0,esk3_0,esk3_0,esk6_0))=k2_yellow_0(esk1_0,k2_enumset1(esk3_0,esk3_0,esk3_0,esk6_0))|~r2_hidden(k2_yellow_0(esk1_0,k2_enumset1(esk3_0,esk3_0,esk3_0,esk6_0)),u1_struct_0(esk2_0))|~m1_subset_1(k2_enumset1(esk3_0,esk3_0,esk3_0,esk6_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_37]), c_0_34])]), c_0_42])).
cnf(c_0_89, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|m1_subset_1(k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_81, c_0_68])).
cnf(c_0_90, negated_conjecture, (r2_hidden(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(esk1_0))|~r2_hidden(esk6_0,u1_struct_0(X1))|~r2_hidden(esk3_0,u1_struct_0(X1))|~m1_yellow_0(X1,esk1_0)|~v5_yellow_0(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_48]), c_0_62]), c_0_35])]), c_0_71])).
cnf(c_0_91, negated_conjecture, (r2_hidden(esk6_0,u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_84, c_0_54])).
cnf(c_0_92, negated_conjecture, (r2_hidden(esk3_0,u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_84, c_0_68])).
cnf(c_0_93, negated_conjecture, (k2_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))!=k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))), inference(rw,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_94, negated_conjecture, (k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0)=k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)|v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_87, c_0_74])).
fof(c_0_95, plain, ![X13]:(v2_struct_0(X13)|~l1_struct_0(X13)|~v1_xboole_0(u1_struct_0(X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_96, negated_conjecture, (k2_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))=k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))|v1_xboole_0(u1_struct_0(esk2_0))|~r2_hidden(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_87]), c_0_89])).
cnf(c_0_97, negated_conjecture, (r2_hidden(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)),u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_34]), c_0_38])]), c_0_92])).
cnf(c_0_98, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))|k2_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))!=k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_99, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_100, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_94]), c_0_97]), c_0_98])).
fof(c_0_101, plain, ![X24]:(~l1_orders_2(X24)|l1_struct_0(X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_102, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~l1_struct_0(esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_42])).
cnf(c_0_103, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_104, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_55])])).
cnf(c_0_105, negated_conjecture, (~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_104]), c_0_71])).
cnf(c_0_106, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_103]), c_0_35])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 107
# Proof object clause steps            : 72
# Proof object formula steps           : 35
# Proof object conjectures             : 57
# Proof object clause conjectures      : 54
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 32
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 33
# Proof object simplifying inferences  : 73
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 48
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 41
# Processed clauses                    : 288
# ...of these trivial                  : 0
# ...subsumed                          : 18
# ...remaining for further processing  : 270
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 30
# Backward-rewritten                   : 53
# Generated clauses                    : 400
# ...of the previous two non-trivial   : 340
# Contextual simplify-reflections      : 15
# Paramodulations                      : 400
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 146
#    Positive orientable unit clauses  : 36
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 106
# Current number of unprocessed clauses: 134
# ...number of literals in the above   : 558
# Current number of archived formulas  : 0
# Current number of archived clauses   : 125
# Clause-clause subsumption calls (NU) : 7847
# Rec. Clause-clause subsumption calls : 1054
# Non-unit clause-clause subsumptions  : 62
# Unit Clause-clause subsumption calls : 59
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 11
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 18511

# -------------------------------------------------
# User time                : 0.023 s
# System time              : 0.000 s
# Total time               : 0.023 s
# Maximum resident set size: 29020 pages
