# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/3a2ced3e5ecd40c7b1d2952e17bb2bde.E_input', t4_xboole_0)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/3a2ced3e5ecd40c7b1d2952e17bb2bde.E_input', t2_boole)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/3a2ced3e5ecd40c7b1d2952e17bb2bde.E_input', t65_xboole_1)).
fof(d1_relat_1, axiom, ![X1]:(v1_relat_1(X1)<=>![X2]:~((r2_hidden(X2,X1)&![X3, X4]:~(X2=k4_tarski(X3,X4))))), file('data/MPTP/proofs/3a2ced3e5ecd40c7b1d2952e17bb2bde.E_input', d1_relat_1)).
fof(d8_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(X3=k5_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>?[X6]:(r2_hidden(k4_tarski(X4,X6),X1)&r2_hidden(k4_tarski(X6,X5),X2))))))), file('data/MPTP/proofs/3a2ced3e5ecd40c7b1d2952e17bb2bde.E_input', d8_relat_1)).
fof(t55_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>k5_relat_1(k5_relat_1(X1,X2),X3)=k5_relat_1(X1,k5_relat_1(X2,X3))))), file('data/MPTP/proofs/3a2ced3e5ecd40c7b1d2952e17bb2bde.E_input', t55_relat_1)).
fof(t62_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>(k5_relat_1(k1_xboole_0,X1)=k1_xboole_0&k5_relat_1(X1,k1_xboole_0)=k1_xboole_0)), file('data/MPTP/proofs/3a2ced3e5ecd40c7b1d2952e17bb2bde.E_input', t62_relat_1)).
fof(dt_k5_relat_1, axiom, ![X1, X2]:((v1_relat_1(X1)&v1_relat_1(X2))=>v1_relat_1(k5_relat_1(X1,X2))), file('data/MPTP/proofs/3a2ced3e5ecd40c7b1d2952e17bb2bde.E_input', dt_k5_relat_1)).
fof(c_0_8, plain, ![X11, X12, X14, X15, X16]:((r1_xboole_0(X11,X12)|r2_hidden(esk2_2(X11,X12),k3_xboole_0(X11,X12)))&(~r2_hidden(X16,k3_xboole_0(X14,X15))|~r1_xboole_0(X14,X15))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_9, plain, ![X10]:k3_xboole_0(X10,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_10, plain, ![X17]:r1_xboole_0(X17,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_11, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_14, plain, ![X18, X19, X22, X24, X25]:((~v1_relat_1(X18)|(~r2_hidden(X19,X18)|X19=k4_tarski(esk3_2(X18,X19),esk4_2(X18,X19))))&((r2_hidden(esk5_1(X22),X22)|v1_relat_1(X22))&(esk5_1(X22)!=k4_tarski(X24,X25)|v1_relat_1(X22)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_relat_1])])])])])])])).
fof(c_0_15, plain, ![X26, X27, X28, X29, X30, X32, X33, X34, X37]:((((r2_hidden(k4_tarski(X29,esk6_5(X26,X27,X28,X29,X30)),X26)|~r2_hidden(k4_tarski(X29,X30),X28)|X28!=k5_relat_1(X26,X27)|~v1_relat_1(X28)|~v1_relat_1(X27)|~v1_relat_1(X26))&(r2_hidden(k4_tarski(esk6_5(X26,X27,X28,X29,X30),X30),X27)|~r2_hidden(k4_tarski(X29,X30),X28)|X28!=k5_relat_1(X26,X27)|~v1_relat_1(X28)|~v1_relat_1(X27)|~v1_relat_1(X26)))&(~r2_hidden(k4_tarski(X32,X34),X26)|~r2_hidden(k4_tarski(X34,X33),X27)|r2_hidden(k4_tarski(X32,X33),X28)|X28!=k5_relat_1(X26,X27)|~v1_relat_1(X28)|~v1_relat_1(X27)|~v1_relat_1(X26)))&((~r2_hidden(k4_tarski(esk7_3(X26,X27,X28),esk8_3(X26,X27,X28)),X28)|(~r2_hidden(k4_tarski(esk7_3(X26,X27,X28),X37),X26)|~r2_hidden(k4_tarski(X37,esk8_3(X26,X27,X28)),X27))|X28=k5_relat_1(X26,X27)|~v1_relat_1(X28)|~v1_relat_1(X27)|~v1_relat_1(X26))&((r2_hidden(k4_tarski(esk7_3(X26,X27,X28),esk9_3(X26,X27,X28)),X26)|r2_hidden(k4_tarski(esk7_3(X26,X27,X28),esk8_3(X26,X27,X28)),X28)|X28=k5_relat_1(X26,X27)|~v1_relat_1(X28)|~v1_relat_1(X27)|~v1_relat_1(X26))&(r2_hidden(k4_tarski(esk9_3(X26,X27,X28),esk8_3(X26,X27,X28)),X27)|r2_hidden(k4_tarski(esk7_3(X26,X27,X28),esk8_3(X26,X27,X28)),X28)|X28=k5_relat_1(X26,X27)|~v1_relat_1(X28)|~v1_relat_1(X27)|~v1_relat_1(X26))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_relat_1])])])])])])).
cnf(c_0_16, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])])).
cnf(c_0_17, plain, (r2_hidden(esk5_1(X1),X1)|v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (r2_hidden(k4_tarski(esk7_3(X1,X2,X3),esk9_3(X1,X2,X3)),X1)|r2_hidden(k4_tarski(esk7_3(X1,X2,X3),esk8_3(X1,X2,X3)),X3)|X3=k5_relat_1(X1,X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, plain, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_20, plain, (r2_hidden(k4_tarski(X1,esk6_5(X2,X3,X4,X1,X5)),X2)|~r2_hidden(k4_tarski(X1,X5),X4)|X4!=k5_relat_1(X2,X3)|~v1_relat_1(X4)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (v1_relat_1(k3_xboole_0(X1,X2))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_11, c_0_17])).
fof(c_0_22, plain, ![X39, X40, X41]:(~v1_relat_1(X39)|(~v1_relat_1(X40)|(~v1_relat_1(X41)|k5_relat_1(k5_relat_1(X39,X40),X41)=k5_relat_1(X39,k5_relat_1(X40,X41))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_relat_1])])])).
cnf(c_0_23, plain, (X1=k5_relat_1(k1_xboole_0,X2)|r2_hidden(k4_tarski(esk7_3(k1_xboole_0,X2,X1),esk8_3(k1_xboole_0,X2,X1)),X1)|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_18]), c_0_19])])).
cnf(c_0_24, plain, (X1!=k5_relat_1(k3_xboole_0(X2,X3),X4)|~r2_hidden(k4_tarski(X5,X6),X1)|~r1_xboole_0(X2,X3)|~v1_relat_1(X1)|~v1_relat_1(X4)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_20]), c_0_21])).
cnf(c_0_25, plain, (k5_relat_1(k5_relat_1(X1,X2),X3)=k5_relat_1(X1,k5_relat_1(X2,X3))|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, plain, (k5_relat_1(k1_xboole_0,X1)=k1_xboole_0|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_23]), c_0_19])])).
cnf(c_0_27, plain, (~r2_hidden(k4_tarski(X1,X2),k5_relat_1(k3_xboole_0(X3,X4),X5))|~r1_xboole_0(X3,X4)|~v1_relat_1(k5_relat_1(k3_xboole_0(X3,X4),X5))|~v1_relat_1(X5)), inference(er,[status(thm)],[c_0_24])).
cnf(c_0_28, plain, (k5_relat_1(k1_xboole_0,k5_relat_1(X1,X2))=k5_relat_1(k1_xboole_0,X2)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_19])])).
fof(c_0_29, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>(k5_relat_1(k1_xboole_0,X1)=k1_xboole_0&k5_relat_1(X1,k1_xboole_0)=k1_xboole_0))), inference(assume_negation,[status(cth)],[t62_relat_1])).
fof(c_0_30, plain, ![X8, X9]:(~v1_relat_1(X8)|~v1_relat_1(X9)|v1_relat_1(k5_relat_1(X8,X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_relat_1])])).
cnf(c_0_31, plain, (~r2_hidden(k4_tarski(X1,X2),k5_relat_1(k1_xboole_0,X3))|~v1_relat_1(k5_relat_1(k1_xboole_0,X3))|~v1_relat_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_12]), c_0_13])])).
cnf(c_0_32, plain, (r2_hidden(k4_tarski(esk6_5(X1,X2,X3,X4,X5),X5),X2)|~r2_hidden(k4_tarski(X4,X5),X3)|X3!=k5_relat_1(X1,X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_33, plain, (k5_relat_1(k1_xboole_0,k1_xboole_0)=k5_relat_1(k1_xboole_0,X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_26]), c_0_19])])).
fof(c_0_34, negated_conjecture, (v1_relat_1(esk1_0)&(k5_relat_1(k1_xboole_0,esk1_0)!=k1_xboole_0|k5_relat_1(esk1_0,k1_xboole_0)!=k1_xboole_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])])).
cnf(c_0_35, plain, (v1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_36, plain, (X1!=k5_relat_1(X2,k5_relat_1(k1_xboole_0,X3))|~r2_hidden(k4_tarski(X4,X5),X1)|~v1_relat_1(k5_relat_1(k1_xboole_0,X3))|~v1_relat_1(X3)|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_37, plain, (k5_relat_1(k1_xboole_0,k1_xboole_0)=k1_xboole_0|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_26, c_0_33])).
cnf(c_0_38, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_39, plain, (v1_relat_1(k5_relat_1(X1,k5_relat_1(X2,X3)))|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_25]), c_0_35])).
cnf(c_0_40, plain, (~r2_hidden(k4_tarski(X1,X2),k5_relat_1(X3,k5_relat_1(k1_xboole_0,X4)))|~v1_relat_1(k5_relat_1(k1_xboole_0,X4))|~v1_relat_1(X4)|~v1_relat_1(X3)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_36]), c_0_35])).
cnf(c_0_41, negated_conjecture, (k5_relat_1(k1_xboole_0,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_42, plain, (v1_relat_1(k5_relat_1(X1,k1_xboole_0))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_26]), c_0_19])])).
cnf(c_0_43, negated_conjecture, (~r2_hidden(k4_tarski(X1,X2),k5_relat_1(X3,k1_xboole_0))|~v1_relat_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_19])])).
cnf(c_0_44, negated_conjecture, (v1_relat_1(k5_relat_1(X1,k1_xboole_0))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_42, c_0_38])).
cnf(c_0_45, negated_conjecture, (k5_relat_1(k1_xboole_0,esk1_0)!=k1_xboole_0|k5_relat_1(esk1_0,k1_xboole_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_46, negated_conjecture, (k5_relat_1(X1,k1_xboole_0)=k5_relat_1(k1_xboole_0,X2)|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_23]), c_0_44])).
cnf(c_0_47, negated_conjecture, (k5_relat_1(k1_xboole_0,esk1_0)!=k1_xboole_0|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_38])]), c_0_26])).
cnf(c_0_48, negated_conjecture, (~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_26]), c_0_38])])).
cnf(c_0_49, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_38, c_0_48]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 50
# Proof object clause steps            : 33
# Proof object formula steps           : 17
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 26
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 572
# ...of these trivial                  : 0
# ...subsumed                          : 415
# ...remaining for further processing  : 157
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 21
# Backward-rewritten                   : 4
# Generated clauses                    : 2695
# ...of the previous two non-trivial   : 2282
# Contextual simplify-reflections      : 29
# Paramodulations                      : 2674
# Factorizations                       : 0
# Equation resolutions                 : 18
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 112
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 107
# Current number of unprocessed clauses: 1614
# ...number of literals in the above   : 9684
# Current number of archived formulas  : 0
# Current number of archived clauses   : 45
# Clause-clause subsumption calls (NU) : 9627
# Rec. Clause-clause subsumption calls : 2772
# Non-unit clause-clause subsumptions  : 399
# Unit Clause-clause subsumption calls : 105
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 50761

# -------------------------------------------------
# User time                : 0.028 s
# System time              : 0.009 s
# Total time               : 0.037 s
# Maximum resident set size: 28544 pages
