# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/3a9455dfb9c4469cb07f44de9b8280f7.E_input', l32_xboole_1)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/3a9455dfb9c4469cb07f44de9b8280f7.E_input', t17_xboole_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/3a9455dfb9c4469cb07f44de9b8280f7.E_input', idempotence_k3_xboole_0)).
fof(d2_xboole_0, axiom, k1_xboole_0=o_0_0_xboole_0, file('data/MPTP/proofs/3a9455dfb9c4469cb07f44de9b8280f7.E_input', d2_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/3a9455dfb9c4469cb07f44de9b8280f7.E_input', t28_xboole_1)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/3a9455dfb9c4469cb07f44de9b8280f7.E_input', t41_xboole_1)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/3a9455dfb9c4469cb07f44de9b8280f7.E_input', t2_boole)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/3a9455dfb9c4469cb07f44de9b8280f7.E_input', commutativity_k3_xboole_0)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/3a9455dfb9c4469cb07f44de9b8280f7.E_input', d7_xboole_0)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/3a9455dfb9c4469cb07f44de9b8280f7.E_input', t4_xboole_0)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/3a9455dfb9c4469cb07f44de9b8280f7.E_input', t3_xboole_0)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/3a9455dfb9c4469cb07f44de9b8280f7.E_input', t1_boole)).
fof(t16_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3)), file('data/MPTP/proofs/3a9455dfb9c4469cb07f44de9b8280f7.E_input', t16_xboole_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/3a9455dfb9c4469cb07f44de9b8280f7.E_input', d5_xboole_0)).
fof(t40_xboole_1, axiom, ![X1, X2]:k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/3a9455dfb9c4469cb07f44de9b8280f7.E_input', t40_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/3a9455dfb9c4469cb07f44de9b8280f7.E_input', commutativity_k2_xboole_0)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/3a9455dfb9c4469cb07f44de9b8280f7.E_input', t36_xboole_1)).
fof(t47_xboole_1, conjecture, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/3a9455dfb9c4469cb07f44de9b8280f7.E_input', t47_xboole_1)).
fof(c_0_18, plain, ![X5, X6]:((k4_xboole_0(X5,X6)!=k1_xboole_0|r1_tarski(X5,X6))&(~r1_tarski(X5,X6)|k4_xboole_0(X5,X6)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
fof(c_0_19, plain, ![X7, X8]:r1_tarski(k3_xboole_0(X7,X8),X7), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_20, plain, ![X9]:k3_xboole_0(X9,X9)=X9, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_21, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, plain, (k1_xboole_0=o_0_0_xboole_0), inference(split_conjunct,[status(thm)],[d2_xboole_0])).
cnf(c_0_23, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_25, plain, ![X10, X11]:(~r1_tarski(X10,X11)|k3_xboole_0(X10,X11)=X10), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_26, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, plain, (k4_xboole_0(X1,X2)=o_0_0_xboole_0|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_29, plain, ![X12, X13, X14]:k4_xboole_0(k4_xboole_0(X12,X13),X14)=k4_xboole_0(X12,k2_xboole_0(X13,X14)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
fof(c_0_30, plain, ![X15]:k3_xboole_0(X15,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_31, plain, ![X16, X17]:k3_xboole_0(X16,X17)=k3_xboole_0(X17,X16), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_32, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=o_0_0_xboole_0), inference(rw,[status(thm)],[c_0_26, c_0_22])).
cnf(c_0_34, plain, (k4_xboole_0(X1,X1)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_35, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, plain, (k3_xboole_0(X1,X2)=X1|k4_xboole_0(X1,X2)!=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_39, plain, (k4_xboole_0(X1,k2_xboole_0(X2,k4_xboole_0(X1,X2)))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40, plain, (k3_xboole_0(X1,o_0_0_xboole_0)=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_22]), c_0_22])).
fof(c_0_41, plain, ![X18, X19]:((~r1_xboole_0(X18,X19)|k3_xboole_0(X18,X19)=k1_xboole_0)&(k3_xboole_0(X18,X19)!=k1_xboole_0|r1_xboole_0(X18,X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
fof(c_0_42, plain, ![X20, X21, X23, X24, X25]:((r1_xboole_0(X20,X21)|r2_hidden(esk1_2(X20,X21),k3_xboole_0(X20,X21)))&(~r2_hidden(X25,k3_xboole_0(X23,X24))|~r1_xboole_0(X23,X24))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_43, plain, ![X26, X27, X29, X30, X31]:(((r2_hidden(esk2_2(X26,X27),X26)|r1_xboole_0(X26,X27))&(r2_hidden(esk2_2(X26,X27),X27)|r1_xboole_0(X26,X27)))&(~r2_hidden(X31,X29)|~r2_hidden(X31,X30)|~r1_xboole_0(X29,X30))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_44, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_23, c_0_37])).
cnf(c_0_45, plain, (k3_xboole_0(X1,k2_xboole_0(X2,k4_xboole_0(X1,X2)))=X1), inference(spm,[status(thm)],[c_0_38, c_0_39])).
fof(c_0_46, plain, ![X32]:k2_xboole_0(X32,k1_xboole_0)=X32, inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_47, plain, (r1_tarski(o_0_0_xboole_0,X1)), inference(spm,[status(thm)],[c_0_23, c_0_40])).
fof(c_0_48, plain, ![X33, X34, X35]:k3_xboole_0(k3_xboole_0(X33,X34),X35)=k3_xboole_0(X33,k3_xboole_0(X34,X35)), inference(variable_rename,[status(thm)],[t16_xboole_1])).
cnf(c_0_49, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_50, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_51, plain, (r2_hidden(esk2_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_52, plain, ![X36, X37, X38, X39, X40, X41, X42, X43]:((((r2_hidden(X39,X36)|~r2_hidden(X39,X38)|X38!=k4_xboole_0(X36,X37))&(~r2_hidden(X39,X37)|~r2_hidden(X39,X38)|X38!=k4_xboole_0(X36,X37)))&(~r2_hidden(X40,X36)|r2_hidden(X40,X37)|r2_hidden(X40,X38)|X38!=k4_xboole_0(X36,X37)))&((~r2_hidden(esk3_3(X41,X42,X43),X43)|(~r2_hidden(esk3_3(X41,X42,X43),X41)|r2_hidden(esk3_3(X41,X42,X43),X42))|X43=k4_xboole_0(X41,X42))&((r2_hidden(esk3_3(X41,X42,X43),X41)|r2_hidden(esk3_3(X41,X42,X43),X43)|X43=k4_xboole_0(X41,X42))&(~r2_hidden(esk3_3(X41,X42,X43),X42)|r2_hidden(esk3_3(X41,X42,X43),X43)|X43=k4_xboole_0(X41,X42))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_53, plain, (r1_tarski(X1,k2_xboole_0(X2,k4_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_54, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_55, plain, ![X45, X46]:k4_xboole_0(k2_xboole_0(X45,X46),X46)=k4_xboole_0(X45,X46), inference(variable_rename,[status(thm)],[t40_xboole_1])).
fof(c_0_56, plain, ![X47, X48]:k2_xboole_0(X47,X48)=k2_xboole_0(X48,X47), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_57, plain, (k4_xboole_0(o_0_0_xboole_0,X1)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_27, c_0_47])).
cnf(c_0_58, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59, plain, (k3_xboole_0(X1,X2)=o_0_0_xboole_0|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_49, c_0_22])).
cnf(c_0_60, plain, (r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_61, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_62, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_63, plain, (r1_tarski(k4_xboole_0(X1,X2),k2_xboole_0(X3,k4_xboole_0(X1,k2_xboole_0(X2,X3))))), inference(spm,[status(thm)],[c_0_53, c_0_35])).
cnf(c_0_64, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X2)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_27, c_0_44])).
cnf(c_0_65, plain, (k2_xboole_0(X1,o_0_0_xboole_0)=X1), inference(rw,[status(thm)],[c_0_54, c_0_22])).
fof(c_0_66, plain, ![X49, X50]:r1_tarski(k4_xboole_0(X49,X50),X49), inference(variable_rename,[status(thm)],[t36_xboole_1])).
cnf(c_0_67, plain, (k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_68, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_69, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_34]), c_0_57])).
cnf(c_0_70, plain, (k3_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_58, c_0_24])).
cnf(c_0_71, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X3))=o_0_0_xboole_0|~r1_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_72, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_61])).
cnf(c_0_73, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_62])).
cnf(c_0_74, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_75, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,k2_xboole_0(X2,X3)),X2),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_65])).
cnf(c_0_76, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_77, plain, (k4_xboole_0(k2_xboole_0(X1,X2),X1)=k4_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_78, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(spm,[status(thm)],[c_0_38, c_0_69])).
cnf(c_0_79, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_70, c_0_37])).
cnf(c_0_80, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X3))=o_0_0_xboole_0|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_71]), c_0_58])).
cnf(c_0_81, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r2_hidden(esk2_2(X1,k4_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_72, c_0_51])).
cnf(c_0_82, plain, (r2_hidden(esk2_2(k4_xboole_0(X1,X2),X3),X1)|r1_xboole_0(k4_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_83, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,k2_xboole_0(X2,X3)),X3),X2)), inference(spm,[status(thm)],[c_0_75, c_0_68])).
cnf(c_0_84, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(X2,X3)),k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_76, c_0_35])).
cnf(c_0_85, plain, (k2_xboole_0(X1,X2)=X1|k4_xboole_0(X2,X1)!=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_77]), c_0_37]), c_0_78])).
cnf(c_0_86, plain, (k3_xboole_0(X1,X2)=o_0_0_xboole_0|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_87, plain, (r1_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X3,X1))), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_88, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_83, c_0_45])).
cnf(c_0_89, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(X2,X3)),k4_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_84, c_0_68])).
cnf(c_0_90, plain, (k2_xboole_0(X1,k3_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_85, c_0_64])).
cnf(c_0_91, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X3))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_92, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X1)))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_88]), c_0_37])).
cnf(c_0_93, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X1)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_27, c_0_23])).
cnf(c_0_94, plain, (k4_xboole_0(k3_xboole_0(X1,k3_xboole_0(X2,X3)),X3)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_64, c_0_58])).
cnf(c_0_95, plain, (r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X3,X2)))), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_96, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,X1)))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_97, plain, (k3_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_76]), c_0_37])).
cnf(c_0_98, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(spm,[status(thm)],[c_0_85, c_0_93])).
cnf(c_0_99, plain, (k3_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,X2)|k4_xboole_0(X1,k2_xboole_0(X2,X3))!=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_38, c_0_35])).
cnf(c_0_100, plain, (k4_xboole_0(k3_xboole_0(X1,X2),k2_xboole_0(X3,k4_xboole_0(X2,X3)))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_94, c_0_45])).
cnf(c_0_101, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X3,X2)))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_32, c_0_95])).
cnf(c_0_102, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_96]), c_0_97])).
cnf(c_0_103, plain, (r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_89, c_0_98])).
cnf(c_0_104, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(k3_xboole_0(X3,X1),X2))=k4_xboole_0(k3_xboole_0(X3,X1),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_37])).
cnf(c_0_105, plain, (k4_xboole_0(X1,k3_xboole_0(X2,k4_xboole_0(X3,X1)))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_97])).
cnf(c_0_106, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X2,X3)))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_32, c_0_103])).
cnf(c_0_107, plain, (~r2_hidden(X1,k3_xboole_0(X2,k3_xboole_0(X3,X4)))|~r1_xboole_0(k3_xboole_0(X2,X3),X4)), inference(spm,[status(thm)],[c_0_50, c_0_58])).
cnf(c_0_108, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk1_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_109, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,X2),X3),k4_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_23, c_0_104])).
cnf(c_0_110, plain, (k4_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X3,X1))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_111, plain, (r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(k3_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_112, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_81, c_0_74])).
cnf(c_0_113, plain, (r1_tarski(k3_xboole_0(X1,X2),k4_xboole_0(X2,k4_xboole_0(X3,X1)))), inference(spm,[status(thm)],[c_0_109, c_0_110])).
cnf(c_0_114, plain, (r1_xboole_0(X1,k3_xboole_0(X2,k4_xboole_0(X3,k3_xboole_0(X1,X2))))), inference(spm,[status(thm)],[c_0_111, c_0_112])).
cnf(c_0_115, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X3,X1)))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_113]), c_0_58]), c_0_97])).
cnf(c_0_116, plain, (r1_xboole_0(X1,k4_xboole_0(X2,k3_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_114, c_0_97])).
fof(c_0_117, negated_conjecture, ~(![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(assume_negation,[status(cth)],[t47_xboole_1])).
cnf(c_0_118, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X2,X1)), inference(rw,[status(thm)],[c_0_92, c_0_115])).
cnf(c_0_119, plain, (r1_xboole_0(X1,k4_xboole_0(X2,k3_xboole_0(X2,X1)))), inference(spm,[status(thm)],[c_0_116, c_0_37])).
fof(c_0_120, negated_conjecture, k4_xboole_0(esk4_0,k3_xboole_0(esk4_0,esk5_0))!=k4_xboole_0(esk4_0,esk5_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_117])])])).
cnf(c_0_121, plain, (k4_xboole_0(X1,k3_xboole_0(X2,X1))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_118]), c_0_37]), c_0_97])).
cnf(c_0_122, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k3_xboole_0(X2,X1)))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_59, c_0_119])).
cnf(c_0_123, negated_conjecture, (k4_xboole_0(esk4_0,k3_xboole_0(esk4_0,esk5_0))!=k4_xboole_0(esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_124, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_35]), c_0_65]), c_0_35]), c_0_68]), c_0_90])).
cnf(c_0_125, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_123, c_0_124])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 126
# Proof object clause steps            : 90
# Proof object formula steps           : 36
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 59
# Proof object simplifying inferences  : 30
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 28
# Processed clauses                    : 7143
# ...of these trivial                  : 1764
# ...subsumed                          : 4400
# ...remaining for further processing  : 979
# Other redundant clauses eliminated   : 578
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 91
# Generated clauses                    : 169563
# ...of the previous two non-trivial   : 78374
# Contextual simplify-reflections      : 1
# Paramodulations                      : 168979
# Factorizations                       : 6
# Equation resolutions                 : 578
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 878
#    Positive orientable unit clauses  : 617
#    Positive unorientable unit clauses: 5
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 254
# Current number of unprocessed clauses: 71015
# ...number of literals in the above   : 116500
# Current number of archived formulas  : 0
# Current number of archived clauses   : 94
# Clause-clause subsumption calls (NU) : 17507
# Rec. Clause-clause subsumption calls : 15461
# Non-unit clause-clause subsumptions  : 1608
# Unit Clause-clause subsumption calls : 224
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3643
# BW rewrite match successes           : 154
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1298834

# -------------------------------------------------
# User time                : 0.750 s
# System time              : 0.043 s
# Total time               : 0.793 s
# Maximum resident set size: 63344 pages
