# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/3a9fc9fcb85f4a93b76621e2d177268d.E_input', d4_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/3a9fc9fcb85f4a93b76621e2d177268d.E_input', d3_tarski)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/3a9fc9fcb85f4a93b76621e2d177268d.E_input', t113_zfmisc_1)).
fof(t123_zfmisc_1, axiom, ![X1, X2, X3, X4]:k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)), file('data/MPTP/proofs/3a9fc9fcb85f4a93b76621e2d177268d.E_input', t123_zfmisc_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/3a9fc9fcb85f4a93b76621e2d177268d.E_input', t3_boole)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/3a9fc9fcb85f4a93b76621e2d177268d.E_input', l32_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/3a9fc9fcb85f4a93b76621e2d177268d.E_input', t28_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/3a9fc9fcb85f4a93b76621e2d177268d.E_input', d10_xboole_0)).
fof(t118_zfmisc_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>(r1_tarski(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&r1_tarski(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2)))), file('data/MPTP/proofs/3a9fc9fcb85f4a93b76621e2d177268d.E_input', t118_zfmisc_1)).
fof(t138_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=>(k2_zfmisc_1(X1,X2)=k1_xboole_0|(r1_tarski(X1,X3)&r1_tarski(X2,X4)))), file('data/MPTP/proofs/3a9fc9fcb85f4a93b76621e2d177268d.E_input', t138_zfmisc_1)).
fof(t117_zfmisc_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&((r1_tarski(k2_zfmisc_1(X2,X1),k2_zfmisc_1(X3,X1))|r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))&~(r1_tarski(X2,X3))))), file('data/MPTP/proofs/3a9fc9fcb85f4a93b76621e2d177268d.E_input', t117_zfmisc_1)).
fof(c_0_11, plain, ![X5, X6, X7, X8, X9, X10, X11, X12]:((((r2_hidden(X8,X5)|~r2_hidden(X8,X7)|X7!=k3_xboole_0(X5,X6))&(r2_hidden(X8,X6)|~r2_hidden(X8,X7)|X7!=k3_xboole_0(X5,X6)))&(~r2_hidden(X9,X5)|~r2_hidden(X9,X6)|r2_hidden(X9,X7)|X7!=k3_xboole_0(X5,X6)))&((~r2_hidden(esk1_3(X10,X11,X12),X12)|(~r2_hidden(esk1_3(X10,X11,X12),X10)|~r2_hidden(esk1_3(X10,X11,X12),X11))|X12=k3_xboole_0(X10,X11))&((r2_hidden(esk1_3(X10,X11,X12),X10)|r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k3_xboole_0(X10,X11))&(r2_hidden(esk1_3(X10,X11,X12),X11)|r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k3_xboole_0(X10,X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_12, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_13, plain, ![X14, X15, X16, X17, X18]:((~r1_tarski(X14,X15)|(~r2_hidden(X16,X14)|r2_hidden(X16,X15)))&((r2_hidden(esk2_2(X17,X18),X17)|r1_tarski(X17,X18))&(~r2_hidden(esk2_2(X17,X18),X18)|r1_tarski(X17,X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_14, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_12])).
cnf(c_0_15, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_16, plain, ![X20, X21]:((k2_zfmisc_1(X20,X21)!=k1_xboole_0|(X20=k1_xboole_0|X21=k1_xboole_0))&((X20!=k1_xboole_0|k2_zfmisc_1(X20,X21)=k1_xboole_0)&(X21!=k1_xboole_0|k2_zfmisc_1(X20,X21)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_17, plain, (r1_tarski(k3_xboole_0(X1,X2),X3)|r2_hidden(esk2_2(k3_xboole_0(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
fof(c_0_18, plain, ![X27, X28, X29, X30]:k2_zfmisc_1(k3_xboole_0(X27,X28),k3_xboole_0(X29,X30))=k3_xboole_0(k2_zfmisc_1(X27,X29),k2_zfmisc_1(X28,X30)), inference(variable_rename,[status(thm)],[t123_zfmisc_1])).
cnf(c_0_19, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_20, plain, ![X22]:k4_xboole_0(X22,k1_xboole_0)=X22, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_21, plain, ![X23, X24]:((k4_xboole_0(X23,X24)!=k1_xboole_0|r1_tarski(X23,X24))&(~r1_tarski(X23,X24)|k4_xboole_0(X23,X24)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_22, plain, (r1_tarski(X1,X2)|~r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23, plain, (r1_tarski(k3_xboole_0(X1,k3_xboole_0(X2,X3)),X4)|r2_hidden(esk2_2(k3_xboole_0(X1,k3_xboole_0(X2,X3)),X4),X3)), inference(spm,[status(thm)],[c_0_14, c_0_17])).
fof(c_0_24, plain, ![X25, X26]:(~r1_tarski(X25,X26)|k3_xboole_0(X25,X26)=X25), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_25, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_19])).
fof(c_0_27, plain, ![X34, X35]:(((r1_tarski(X34,X35)|X34!=X35)&(r1_tarski(X35,X34)|X34!=X35))&(~r1_tarski(X34,X35)|~r1_tarski(X35,X34)|X34=X35)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_28, plain, ![X31, X32, X33]:((r1_tarski(k2_zfmisc_1(X31,X33),k2_zfmisc_1(X32,X33))|~r1_tarski(X31,X32))&(r1_tarski(k2_zfmisc_1(X33,X31),k2_zfmisc_1(X33,X32))|~r1_tarski(X31,X32))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t118_zfmisc_1])])])).
cnf(c_0_29, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_30, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_31, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_32, plain, (r1_tarski(k3_xboole_0(X1,k3_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_33, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k1_xboole_0)=k2_zfmisc_1(k3_xboole_0(X1,X3),k3_xboole_0(X2,k1_xboole_0))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_35, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_29])).
fof(c_0_38, negated_conjecture, ~(![X1, X2, X3, X4]:(r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=>(k2_zfmisc_1(X1,X2)=k1_xboole_0|(r1_tarski(X1,X3)&r1_tarski(X2,X4))))), inference(assume_negation,[status(cth)],[t138_zfmisc_1])).
cnf(c_0_39, plain, (k1_xboole_0=X1|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_40, plain, (r1_tarski(k3_xboole_0(X1,X2),X3)|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_41, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_42, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k1_xboole_0)=k2_zfmisc_1(X1,k3_xboole_0(X2,k1_xboole_0))|~r1_tarski(X1,X3)), inference(spm,[status(thm)],[c_0_34, c_0_33])).
cnf(c_0_43, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_35])).
cnf(c_0_44, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k1_xboole_0)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_45, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k2_zfmisc_1(X1,X3)|~r1_tarski(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(spm,[status(thm)],[c_0_25, c_0_33])).
fof(c_0_46, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk5_0,esk6_0))&(k2_zfmisc_1(esk3_0,esk4_0)!=k1_xboole_0&(~r1_tarski(esk3_0,esk5_0)|~r1_tarski(esk4_0,esk6_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_38])])])).
cnf(c_0_47, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_48, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_49, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_50, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k1_xboole_0)=k2_zfmisc_1(X1,k3_xboole_0(X2,k1_xboole_0))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_51, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k1_xboole_0)|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))|~r1_tarski(k3_xboole_0(X1,X3),k1_xboole_0)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_52, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk5_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_53, plain, (k1_xboole_0=X1|~r1_tarski(X2,k1_xboole_0)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_33, c_0_47])).
cnf(c_0_54, plain, (r1_tarski(X1,k1_xboole_0)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_48, c_0_30])).
cnf(c_0_55, plain, (r1_tarski(k3_xboole_0(X1,X2),X3)|r2_hidden(esk2_2(k3_xboole_0(X1,X2),X3),X1)), inference(spm,[status(thm)],[c_0_49, c_0_15])).
cnf(c_0_56, plain, (r2_hidden(X1,k1_xboole_0)|X2!=k2_zfmisc_1(X3,k3_xboole_0(X4,k1_xboole_0))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_12, c_0_50])).
cnf(c_0_57, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k1_xboole_0)|~r1_tarski(k3_xboole_0(esk3_0,esk5_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_58, plain, (k1_xboole_0=X1|X2!=k1_xboole_0|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_59, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_22, c_0_55])).
cnf(c_0_60, plain, (r2_hidden(X1,k1_xboole_0)|X2!=k2_zfmisc_1(X3,X4)|~r1_tarski(X4,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_56, c_0_33])).
cnf(c_0_61, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k1_xboole_0)|k3_xboole_0(esk3_0,esk5_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_57, c_0_54])).
cnf(c_0_62, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_63, plain, (r2_hidden(X1,k1_xboole_0)|X2!=k2_zfmisc_1(X3,X4)|X4!=k1_xboole_0|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_60, c_0_54])).
fof(c_0_64, plain, ![X40, X41, X42]:((~r1_tarski(k2_zfmisc_1(X41,X40),k2_zfmisc_1(X42,X40))|r1_tarski(X41,X42)|X40=k1_xboole_0)&(~r1_tarski(k2_zfmisc_1(X40,X41),k2_zfmisc_1(X40,X42))|r1_tarski(X41,X42)|X40=k1_xboole_0)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t117_zfmisc_1])])])])).
cnf(c_0_65, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,k3_xboole_0(X2,X4)))|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X5,X4))|~r1_tarski(k3_xboole_0(X1,X5),X3)), inference(spm,[status(thm)],[c_0_36, c_0_45])).
cnf(c_0_66, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k1_xboole_0)|esk3_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_67, negated_conjecture, (k2_zfmisc_1(esk3_0,esk4_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_68, plain, (r2_hidden(X1,k1_xboole_0)|X2!=k1_xboole_0|~r2_hidden(X1,k2_zfmisc_1(X3,X2))), inference(er,[status(thm)],[c_0_63])).
cnf(c_0_69, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_22, c_0_17])).
cnf(c_0_70, plain, (r1_tarski(X2,X3)|X1=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_71, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(X1,k3_xboole_0(esk4_0,esk6_0)))|~r1_tarski(k3_xboole_0(esk3_0,esk5_0),X1)), inference(spm,[status(thm)],[c_0_65, c_0_52])).
cnf(c_0_72, negated_conjecture, (esk3_0!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_66]), c_0_67])).
cnf(c_0_73, plain, (r1_tarski(k3_xboole_0(X1,k2_zfmisc_1(X2,X3)),X4)|r2_hidden(esk2_2(k3_xboole_0(X1,k2_zfmisc_1(X2,X3)),X4),k1_xboole_0)|X3!=k1_xboole_0), inference(spm,[status(thm)],[c_0_68, c_0_17])).
cnf(c_0_74, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_39, c_0_69])).
cnf(c_0_75, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_76, negated_conjecture, (r1_tarski(esk4_0,k3_xboole_0(esk4_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_59])]), c_0_72])).
cnf(c_0_77, plain, (r1_tarski(k3_xboole_0(X1,k2_zfmisc_1(X2,X3)),k1_xboole_0)|X3!=k1_xboole_0), inference(spm,[status(thm)],[c_0_22, c_0_73])).
cnf(c_0_78, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_59, c_0_74])).
cnf(c_0_79, negated_conjecture, (k3_xboole_0(esk4_0,esk6_0)=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_59])])).
cnf(c_0_80, plain, (k3_xboole_0(X1,k2_zfmisc_1(X2,X3))=k1_xboole_0|X3!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_77]), c_0_78])])).
cnf(c_0_81, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(X1,esk4_0))|~r1_tarski(k3_xboole_0(esk3_0,esk5_0),X1)), inference(rw,[status(thm)],[c_0_71, c_0_79])).
cnf(c_0_82, negated_conjecture, (~r1_tarski(esk3_0,esk5_0)|~r1_tarski(esk4_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_83, negated_conjecture, (r1_tarski(esk4_0,esk6_0)), inference(spm,[status(thm)],[c_0_69, c_0_79])).
cnf(c_0_84, plain, (k1_xboole_0=X1|X2!=k1_xboole_0|~r1_tarski(X1,k2_zfmisc_1(X3,X2))), inference(spm,[status(thm)],[c_0_33, c_0_80])).
cnf(c_0_85, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk5_0,esk4_0))), inference(spm,[status(thm)],[c_0_81, c_0_69])).
cnf(c_0_86, plain, (r1_tarski(X1,X3)|X2=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_87, negated_conjecture, (~r1_tarski(esk3_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_83])])).
cnf(c_0_88, negated_conjecture, (esk4_0!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_67])).
cnf(c_0_89, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_85]), c_0_87]), c_0_88]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 90
# Proof object clause steps            : 67
# Proof object formula steps           : 23
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 45
# Proof object simplifying inferences  : 15
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 15217
# ...of these trivial                  : 201
# ...subsumed                          : 11582
# ...remaining for further processing  : 3433
# Other redundant clauses eliminated   : 11
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 966
# Backward-rewritten                   : 204
# Generated clauses                    : 291254
# ...of the previous two non-trivial   : 281105
# Contextual simplify-reflections      : 407
# Paramodulations                      : 290956
# Factorizations                       : 210
# Equation resolutions                 : 88
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2235
#    Positive orientable unit clauses  : 33
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 11
#    Non-unit-clauses                  : 2191
# Current number of unprocessed clauses: 262306
# ...number of literals in the above   : 1490467
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1196
# Clause-clause subsumption calls (NU) : 842910
# Rec. Clause-clause subsumption calls : 180993
# Non-unit clause-clause subsumptions  : 10177
# Unit Clause-clause subsumption calls : 9648
# Rewrite failures with RHS unbound    : 88
# BW rewrite match attempts            : 189
# BW rewrite match successes           : 58
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6125205

# -------------------------------------------------
# User time                : 4.543 s
# System time              : 0.096 s
# Total time               : 4.639 s
# Maximum resident set size: 207748 pages
