# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t39_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2))), file('data/MPTP/proofs/3b16289c22174fd99b8438e29d19774b.E_input', t39_yellow_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/3b16289c22174fd99b8438e29d19774b.E_input', t24_orders_2)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/3b16289c22174fd99b8438e29d19774b.E_input', d1_tarski)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/3b16289c22174fd99b8438e29d19774b.E_input', t69_enumset1)).
fof(t63_subset_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))), file('data/MPTP/proofs/3b16289c22174fd99b8438e29d19774b.E_input', t63_subset_1)).
fof(t19_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))))), file('data/MPTP/proofs/3b16289c22174fd99b8438e29d19774b.E_input', t19_yellow_0)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/3b16289c22174fd99b8438e29d19774b.E_input', redefinition_k6_domain_1)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/3b16289c22174fd99b8438e29d19774b.E_input', t5_subset)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/3b16289c22174fd99b8438e29d19774b.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/3b16289c22174fd99b8438e29d19774b.E_input', dt_l1_orders_2)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/3b16289c22174fd99b8438e29d19774b.E_input', t2_subset)).
fof(d10_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_yellow_0(X1,X2)=>(X3=k2_yellow_0(X1,X2)<=>(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))))))), file('data/MPTP/proofs/3b16289c22174fd99b8438e29d19774b.E_input', d10_yellow_0)).
fof(dt_k2_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/3b16289c22174fd99b8438e29d19774b.E_input', dt_k2_yellow_0)).
fof(t7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2)))))))), file('data/MPTP/proofs/3b16289c22174fd99b8438e29d19774b.E_input', t7_yellow_0)).
fof(t30_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4)))))&((r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))=>(X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3)))))), file('data/MPTP/proofs/3b16289c22174fd99b8438e29d19774b.E_input', t30_yellow_0)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/3b16289c22174fd99b8438e29d19774b.E_input', d9_lattice3)).
fof(c_0_16, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))), introduced(definition)).
fof(c_0_17, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))), inference(split_equiv,[status(thm)],[c_0_16])).
fof(c_0_18, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2)))), inference(assume_negation,[status(cth)],[t39_yellow_0])).
fof(c_0_19, plain, ![X53, X54, X55, X56, X57]:(((r1_orders_2(X55,X56,X54)|(X56!=k11_lattice3(X55,X54,X53)|~r2_yellow_0(X55,k2_tarski(X54,X53)))|~m1_subset_1(X56,u1_struct_0(X55))|~epred1_3(X55,X54,X53))&((r1_orders_2(X55,X56,X53)|(X56!=k11_lattice3(X55,X54,X53)|~r2_yellow_0(X55,k2_tarski(X54,X53)))|~m1_subset_1(X56,u1_struct_0(X55))|~epred1_3(X55,X54,X53))&(~m1_subset_1(X57,u1_struct_0(X55))|(~r1_orders_2(X55,X57,X54)|~r1_orders_2(X55,X57,X53)|r1_orders_2(X55,X57,X56))|(X56!=k11_lattice3(X55,X54,X53)|~r2_yellow_0(X55,k2_tarski(X54,X53)))|~m1_subset_1(X56,u1_struct_0(X55))|~epred1_3(X55,X54,X53))))&(((X56=k11_lattice3(X55,X54,X53)|(m1_subset_1(esk7_4(X53,X54,X55,X56),u1_struct_0(X55))|~r1_orders_2(X55,X56,X53)|~r1_orders_2(X55,X56,X54))|~m1_subset_1(X56,u1_struct_0(X55))|~epred1_3(X55,X54,X53))&(r2_yellow_0(X55,k2_tarski(X54,X53))|(m1_subset_1(esk7_4(X53,X54,X55,X56),u1_struct_0(X55))|~r1_orders_2(X55,X56,X53)|~r1_orders_2(X55,X56,X54))|~m1_subset_1(X56,u1_struct_0(X55))|~epred1_3(X55,X54,X53)))&((((X56=k11_lattice3(X55,X54,X53)|(r1_orders_2(X55,esk7_4(X53,X54,X55,X56),X54)|~r1_orders_2(X55,X56,X53)|~r1_orders_2(X55,X56,X54))|~m1_subset_1(X56,u1_struct_0(X55))|~epred1_3(X55,X54,X53))&(r2_yellow_0(X55,k2_tarski(X54,X53))|(r1_orders_2(X55,esk7_4(X53,X54,X55,X56),X54)|~r1_orders_2(X55,X56,X53)|~r1_orders_2(X55,X56,X54))|~m1_subset_1(X56,u1_struct_0(X55))|~epred1_3(X55,X54,X53)))&((X56=k11_lattice3(X55,X54,X53)|(r1_orders_2(X55,esk7_4(X53,X54,X55,X56),X53)|~r1_orders_2(X55,X56,X53)|~r1_orders_2(X55,X56,X54))|~m1_subset_1(X56,u1_struct_0(X55))|~epred1_3(X55,X54,X53))&(r2_yellow_0(X55,k2_tarski(X54,X53))|(r1_orders_2(X55,esk7_4(X53,X54,X55,X56),X53)|~r1_orders_2(X55,X56,X53)|~r1_orders_2(X55,X56,X54))|~m1_subset_1(X56,u1_struct_0(X55))|~epred1_3(X55,X54,X53))))&((X56=k11_lattice3(X55,X54,X53)|(~r1_orders_2(X55,esk7_4(X53,X54,X55,X56),X56)|~r1_orders_2(X55,X56,X53)|~r1_orders_2(X55,X56,X54))|~m1_subset_1(X56,u1_struct_0(X55))|~epred1_3(X55,X54,X53))&(r2_yellow_0(X55,k2_tarski(X54,X53))|(~r1_orders_2(X55,esk7_4(X53,X54,X55,X56),X56)|~r1_orders_2(X55,X56,X53)|~r1_orders_2(X55,X56,X54))|~m1_subset_1(X56,u1_struct_0(X55))|~epred1_3(X55,X54,X53)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])])])).
fof(c_0_20, plain, ![X44, X45]:(v2_struct_0(X44)|(~v3_orders_2(X44)|~l1_orders_2(X44))|(~m1_subset_1(X45,u1_struct_0(X44))|r1_orders_2(X44,X45,X45))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_21, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_18])])])])).
cnf(c_0_22, plain, (X1=k11_lattice3(X2,X3,X4)|~r1_orders_2(X2,esk7_4(X4,X3,X2,X1),X1)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23, plain, (X1=k11_lattice3(X2,X3,X4)|r1_orders_2(X2,esk7_4(X4,X3,X2,X1),X4)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_26, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (X1=k11_lattice3(X2,X3,X4)|r1_orders_2(X2,esk7_4(X4,X3,X2,X1),X3)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_29, plain, ![X19, X20, X21, X22, X23, X24]:(((~r2_hidden(X21,X20)|X21=X19|X20!=k1_tarski(X19))&(X22!=X19|r2_hidden(X22,X20)|X20!=k1_tarski(X19)))&((~r2_hidden(esk5_2(X23,X24),X24)|esk5_2(X23,X24)!=X23|X24=k1_tarski(X23))&(r2_hidden(esk5_2(X23,X24),X24)|esk5_2(X23,X24)=X23|X24=k1_tarski(X23)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_30, plain, ![X27]:k2_tarski(X27,X27)=k1_tarski(X27), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_31, plain, ![X51, X52]:(~r2_hidden(X51,X52)|m1_subset_1(k1_tarski(X51),k1_zfmisc_1(X52))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_subset_1])])).
cnf(c_0_32, plain, (k11_lattice3(X1,X2,X3)=X3|~epred1_3(X1,X2,X3)|~r1_orders_2(X1,X3,X3)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_33, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_34, plain, (k11_lattice3(X1,X2,X3)=X2|~epred1_3(X1,X2,X3)|~r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_22, c_0_28])).
fof(c_0_35, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>epred1_3(X1,X2,X3)))), inference(apply_def,[status(thm)],[t19_yellow_0, c_0_16])).
cnf(c_0_36, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_38, plain, ![X30, X31]:(v1_xboole_0(X30)|~m1_subset_1(X31,X30)|k6_domain_1(X30,X31)=k1_tarski(X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_39, plain, ![X46, X47, X48]:(~r2_hidden(X46,X47)|(~m1_subset_1(X47,k1_zfmisc_1(X48))|~v1_xboole_0(X48))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
cnf(c_0_40, plain, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_42, plain, ![X40]:(v2_struct_0(X40)|~l1_struct_0(X40)|~v1_xboole_0(u1_struct_0(X40))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_43, plain, ![X26]:(~l1_orders_2(X26)|l1_struct_0(X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
fof(c_0_44, plain, ![X49, X50]:(~m1_subset_1(X49,X50)|(v1_xboole_0(X50)|r2_hidden(X49,X50))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_45, negated_conjecture, (k11_lattice3(esk1_0,X1,X2)=X2|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_46, negated_conjecture, (k11_lattice3(esk1_0,X1,X2)=X1|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_34, c_0_33])).
fof(c_0_47, plain, ![X32, X33, X34]:(~v5_orders_2(X32)|~l1_orders_2(X32)|(~m1_subset_1(X33,u1_struct_0(X32))|(~m1_subset_1(X34,u1_struct_0(X32))|epred1_3(X32,X33,X34)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_35])])])).
fof(c_0_48, plain, ![X8, X9, X10, X11]:(((r1_lattice3(X8,X9,X10)|X10!=k2_yellow_0(X8,X9)|~r2_yellow_0(X8,X9)|~m1_subset_1(X10,u1_struct_0(X8))|~l1_orders_2(X8))&(~m1_subset_1(X11,u1_struct_0(X8))|(~r1_lattice3(X8,X9,X11)|r1_orders_2(X8,X11,X10))|X10!=k2_yellow_0(X8,X9)|~r2_yellow_0(X8,X9)|~m1_subset_1(X10,u1_struct_0(X8))|~l1_orders_2(X8)))&((m1_subset_1(esk3_3(X8,X9,X10),u1_struct_0(X8))|~r1_lattice3(X8,X9,X10)|X10=k2_yellow_0(X8,X9)|~r2_yellow_0(X8,X9)|~m1_subset_1(X10,u1_struct_0(X8))|~l1_orders_2(X8))&((r1_lattice3(X8,X9,esk3_3(X8,X9,X10))|~r1_lattice3(X8,X9,X10)|X10=k2_yellow_0(X8,X9)|~r2_yellow_0(X8,X9)|~m1_subset_1(X10,u1_struct_0(X8))|~l1_orders_2(X8))&(~r1_orders_2(X8,esk3_3(X8,X9,X10),X10)|~r1_lattice3(X8,X9,X10)|X10=k2_yellow_0(X8,X9)|~r2_yellow_0(X8,X9)|~m1_subset_1(X10,u1_struct_0(X8))|~l1_orders_2(X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_yellow_0])])])])])).
fof(c_0_49, plain, ![X28, X29]:(~l1_orders_2(X28)|m1_subset_1(k2_yellow_0(X28,X29),u1_struct_0(X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_yellow_0])])).
fof(c_0_50, plain, ![X41, X42, X43]:((~r1_lattice3(X41,k1_tarski(X43),X42)|r1_orders_2(X41,X42,X43)|~m1_subset_1(X43,u1_struct_0(X41))|~m1_subset_1(X42,u1_struct_0(X41))|~l1_orders_2(X41))&((~r1_orders_2(X41,X42,X43)|r1_lattice3(X41,k1_tarski(X43),X42)|~m1_subset_1(X43,u1_struct_0(X41))|~m1_subset_1(X42,u1_struct_0(X41))|~l1_orders_2(X41))&((~r2_lattice3(X41,k1_tarski(X43),X42)|r1_orders_2(X41,X43,X42)|~m1_subset_1(X43,u1_struct_0(X41))|~m1_subset_1(X42,u1_struct_0(X41))|~l1_orders_2(X41))&(~r1_orders_2(X41,X43,X42)|r2_lattice3(X41,k1_tarski(X43),X42)|~m1_subset_1(X43,u1_struct_0(X41))|~m1_subset_1(X42,u1_struct_0(X41))|~l1_orders_2(X41))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])).
cnf(c_0_51, plain, (X1=X3|X2!=k2_tarski(X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_52, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_53, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_54, plain, (m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(X2))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_40, c_0_37])).
cnf(c_0_55, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X2)), inference(rw,[status(thm)],[c_0_41, c_0_37])).
cnf(c_0_56, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_57, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_58, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_59, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_60, negated_conjecture, (X1=X2|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,X1)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_61, plain, (epred1_3(X1,X2,X3)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_62, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_63, plain, (r1_orders_2(X2,X1,X4)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|X4!=k2_yellow_0(X2,X3)|~r2_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_64, plain, (m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_65, plain, (r1_orders_2(X1,X3,X2)|~r1_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_66, plain, (r1_lattice3(X1,X2,X3)|X3!=k2_yellow_0(X1,X2)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_67, plain, (X1=X2|~r2_hidden(X1,k2_tarski(X2,X2))), inference(er,[status(thm)],[c_0_51])).
cnf(c_0_68, plain, (k2_tarski(X2,X2)=k6_domain_1(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_52, c_0_37])).
cnf(c_0_69, plain, (r2_hidden(esk5_2(X1,X2),X2)|esk5_2(X1,X2)=X1|X2=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_70, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k2_tarski(X3,X3))|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_71, plain, (r2_hidden(X1,k2_tarski(X1,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_55])])).
cnf(c_0_72, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_73, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_74, negated_conjecture, (X1=X2|~r1_orders_2(esk1_0,X2,X1)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_26]), c_0_62])])).
cnf(c_0_75, plain, (r1_orders_2(X1,X2,k2_yellow_0(X1,X3))|~r1_lattice3(X1,X3,X2)|~r2_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_63]), c_0_64])).
cnf(c_0_76, plain, (r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,k2_tarski(X2,X2),X3)), inference(rw,[status(thm)],[c_0_65, c_0_37])).
cnf(c_0_77, plain, (r1_lattice3(X1,X2,k2_yellow_0(X1,X2))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_66]), c_0_64])).
cnf(c_0_78, plain, (X1=X2|v1_xboole_0(X3)|~r2_hidden(X1,k6_domain_1(X3,X2))|~m1_subset_1(X2,X3)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_79, plain, (X2=k2_tarski(X1,X1)|esk5_2(X1,X2)=X1|r2_hidden(esk5_2(X1,X2),X2)), inference(rw,[status(thm)],[c_0_69, c_0_37])).
cnf(c_0_80, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_81, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_26])]), c_0_27])).
cnf(c_0_82, negated_conjecture, (k2_yellow_0(esk1_0,X1)=X2|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),X2)|~r1_lattice3(esk1_0,X1,X2)|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_26])])).
cnf(c_0_83, plain, (r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(X2,X2)),X2)|~r2_yellow_0(X1,k2_tarski(X2,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_64])).
cnf(c_0_84, plain, (esk5_2(X1,k6_domain_1(X2,X3))=X1|k6_domain_1(X2,X3)=k2_tarski(X1,X1)|esk5_2(X1,k6_domain_1(X2,X3))=X3|v1_xboole_0(X2)|~m1_subset_1(X3,X2)), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_85, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_86, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|~r1_orders_2(X1,esk7_4(X3,X2,X1,X4),X4)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~epred1_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_87, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|r1_orders_2(X1,esk7_4(X3,X2,X1,X4),X3)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~epred1_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_88, negated_conjecture, (k2_yellow_0(esk1_0,k2_tarski(X1,X1))=X1|~r1_lattice3(esk1_0,k2_tarski(X1,X1),X1)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(X1,X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_26])])).
cnf(c_0_89, plain, (r1_lattice3(X1,k1_tarski(X3),X2)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_90, plain, (X2=k1_tarski(X1)|~r2_hidden(esk5_2(X1,X2),X2)|esk5_2(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_91, negated_conjecture, (esk5_2(X1,k6_domain_1(u1_struct_0(esk1_0),esk2_0))=esk2_0|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(X1,X1)|esk5_2(X1,k6_domain_1(u1_struct_0(esk1_0),esk2_0))=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_59]), c_0_85])).
cnf(c_0_92, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|~epred1_3(X1,X2,X3)|~r1_orders_2(X1,X3,X3)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_93, negated_conjecture, (k2_yellow_0(esk1_0,k2_tarski(X1,X1))=X1|~r1_lattice3(esk1_0,k2_tarski(X1,X1),X1)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_64]), c_0_26])])).
cnf(c_0_94, plain, (r1_lattice3(X1,k2_tarski(X3,X3),X2)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[c_0_89, c_0_37])).
cnf(c_0_95, plain, (X2=k2_tarski(X1,X1)|esk5_2(X1,X2)!=X1|~r2_hidden(esk5_2(X1,X2),X2)), inference(rw,[status(thm)],[c_0_90, c_0_37])).
cnf(c_0_96, negated_conjecture, (esk5_2(esk2_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))=esk2_0|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)), inference(er,[status(thm)],[inference(ef,[status(thm)],[c_0_91])])).
cnf(c_0_97, negated_conjecture, (r2_yellow_0(esk1_0,k2_tarski(X1,X2))|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_92, c_0_33])).
cnf(c_0_98, negated_conjecture, (k2_yellow_0(esk1_0,k2_tarski(X1,X1))=X1|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_26])]), c_0_33])).
cnf(c_0_99, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)|~r2_hidden(esk2_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))), inference(spm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_100, plain, (v1_xboole_0(X1)|r2_hidden(X2,k6_domain_1(X1,X2))|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_71, c_0_68])).
cnf(c_0_101, negated_conjecture, (v1_xboole_0(X1)|r2_yellow_0(esk1_0,k6_domain_1(X1,X2))|~epred1_3(esk1_0,X2,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X2,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_68]), c_0_33])).
cnf(c_0_102, plain, (r1_orders_2(X1,X2,X3)|~r2_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
fof(c_0_103, plain, ![X13, X14, X15, X16, X17]:(((r2_lattice3(X13,X15,X14)|(X14!=k1_yellow_0(X13,X15)|~r1_yellow_0(X13,X15))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13)))&(~m1_subset_1(X16,u1_struct_0(X13))|(~r2_lattice3(X13,X15,X16)|r1_orders_2(X13,X14,X16))|(X14!=k1_yellow_0(X13,X15)|~r1_yellow_0(X13,X15))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13))))&(((X14=k1_yellow_0(X13,X17)|(m1_subset_1(esk4_3(X13,X14,X17),u1_struct_0(X13))|~r2_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13)))&(r1_yellow_0(X13,X17)|(m1_subset_1(esk4_3(X13,X14,X17),u1_struct_0(X13))|~r2_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13))))&(((X14=k1_yellow_0(X13,X17)|(r2_lattice3(X13,X17,esk4_3(X13,X14,X17))|~r2_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13)))&(r1_yellow_0(X13,X17)|(r2_lattice3(X13,X17,esk4_3(X13,X14,X17))|~r2_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13))))&((X14=k1_yellow_0(X13,X17)|(~r1_orders_2(X13,X14,esk4_3(X13,X14,X17))|~r2_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13)))&(r1_yellow_0(X13,X17)|(~r1_orders_2(X13,X14,esk4_3(X13,X14,X17))|~r2_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_yellow_0])])])])])])).
fof(c_0_104, plain, ![X35, X36, X37, X38]:((~r2_lattice3(X35,X36,X37)|(~m1_subset_1(X38,u1_struct_0(X35))|(~r2_hidden(X38,X36)|r1_orders_2(X35,X38,X37)))|~m1_subset_1(X37,u1_struct_0(X35))|~l1_orders_2(X35))&((m1_subset_1(esk6_3(X35,X36,X37),u1_struct_0(X35))|r2_lattice3(X35,X36,X37)|~m1_subset_1(X37,u1_struct_0(X35))|~l1_orders_2(X35))&((r2_hidden(esk6_3(X35,X36,X37),X36)|r2_lattice3(X35,X36,X37)|~m1_subset_1(X37,u1_struct_0(X35))|~l1_orders_2(X35))&(~r1_orders_2(X35,esk6_3(X35,X36,X37),X37)|r2_lattice3(X35,X36,X37)|~m1_subset_1(X37,u1_struct_0(X35))|~l1_orders_2(X35))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
cnf(c_0_105, negated_conjecture, (k2_yellow_0(esk1_0,k6_domain_1(X1,X2))=X2|v1_xboole_0(X1)|~r2_yellow_0(esk1_0,k6_domain_1(X1,X2))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_98, c_0_68])).
cnf(c_0_106, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_59])]), c_0_85])).
cnf(c_0_107, negated_conjecture, (v1_xboole_0(X1)|r2_yellow_0(esk1_0,k6_domain_1(X1,X2))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_61]), c_0_26]), c_0_62])])).
cnf(c_0_108, plain, (r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r2_lattice3(X1,k2_tarski(X2,X2),X3)), inference(rw,[status(thm)],[c_0_102, c_0_37])).
cnf(c_0_109, plain, (X1=k1_yellow_0(X2,X3)|r2_lattice3(X2,X3,esk4_3(X2,X1,X3))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_110, plain, (X1=k1_yellow_0(X2,X3)|m1_subset_1(esk4_3(X2,X1,X3),u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_111, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_112, negated_conjecture, (k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_113, negated_conjecture, (k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))=esk2_0|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_59])]), c_0_85])).
cnf(c_0_114, negated_conjecture, (r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_106]), c_0_59])]), c_0_85])).
cnf(c_0_115, plain, (X1=k1_yellow_0(X2,X3)|~r1_orders_2(X2,X1,esk4_3(X2,X1,X3))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_116, plain, (X1=k1_yellow_0(X2,k2_tarski(X3,X3))|r1_orders_2(X2,X3,esk4_3(X2,X1,k2_tarski(X3,X3)))|~r2_lattice3(X2,k2_tarski(X3,X3),X1)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_110])).
cnf(c_0_117, plain, (r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk6_3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_118, plain, (esk6_3(X1,k2_tarski(X2,X2),X3)=X2|r2_lattice3(X1,k2_tarski(X2,X2),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_67, c_0_111])).
cnf(c_0_119, negated_conjecture, (k1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))!=esk2_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_68]), c_0_59])]), c_0_85])).
cnf(c_0_120, negated_conjecture, (k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_113, c_0_114])])).
cnf(c_0_121, plain, (k1_yellow_0(X1,k2_tarski(X2,X2))=X2|~r2_lattice3(X1,k2_tarski(X2,X2),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_115, c_0_116])).
cnf(c_0_122, plain, (r2_lattice3(X1,k2_tarski(X2,X2),X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_117, c_0_118])).
cnf(c_0_123, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_113]), c_0_59]), c_0_26])])).
cnf(c_0_124, negated_conjecture, (k1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_119, c_0_120])])).
cnf(c_0_125, plain, (k1_yellow_0(X1,k2_tarski(X2,X2))=X2|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_121, c_0_122])).
cnf(c_0_126, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_123, c_0_114])])).
cnf(c_0_127, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125]), c_0_126]), c_0_59]), c_0_26]), c_0_62])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 128
# Proof object clause steps            : 91
# Proof object formula steps           : 37
# Proof object conjectures             : 37
# Proof object clause conjectures      : 34
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 35
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 40
# Proof object simplifying inferences  : 68
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 52
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 51
# Processed clauses                    : 1764
# ...of these trivial                  : 7
# ...subsumed                          : 927
# ...remaining for further processing  : 830
# Other redundant clauses eliminated   : 281
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 266
# Backward-rewritten                   : 82
# Generated clauses                    : 28884
# ...of the previous two non-trivial   : 27650
# Contextual simplify-reflections      : 50
# Paramodulations                      : 28531
# Factorizations                       : 72
# Equation resolutions                 : 282
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 422
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 406
# Current number of unprocessed clauses: 25762
# ...number of literals in the above   : 208250
# Current number of archived formulas  : 0
# Current number of archived clauses   : 400
# Clause-clause subsumption calls (NU) : 57567
# Rec. Clause-clause subsumption calls : 7828
# Non-unit clause-clause subsumptions  : 1117
# Unit Clause-clause subsumption calls : 455
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 13
# BW rewrite match successes           : 7
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 829577

# -------------------------------------------------
# User time                : 0.614 s
# System time              : 0.013 s
# Total time               : 0.627 s
# Maximum resident set size: 28720 pages
