# No SInE strategy applied
# Auto-Mode selected heuristic G_E___200_B02_F1_SE_CS_SP_PI_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t134_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4)))), file('data/MPTP/proofs/3b6aeabae1e044a2927af77a83d131ba.E_input', t134_zfmisc_1)).
fof(t117_zfmisc_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&((r1_tarski(k2_zfmisc_1(X2,X1),k2_zfmisc_1(X3,X1))|r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))&~(r1_tarski(X2,X3))))), file('data/MPTP/proofs/3b6aeabae1e044a2927af77a83d131ba.E_input', t117_zfmisc_1)).
fof(t125_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k4_xboole_0(X1,X2),X3)=k4_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&k2_zfmisc_1(X3,k4_xboole_0(X1,X2))=k4_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))), file('data/MPTP/proofs/3b6aeabae1e044a2927af77a83d131ba.E_input', t125_zfmisc_1)).
fof(t123_zfmisc_1, axiom, ![X1, X2, X3, X4]:k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)), file('data/MPTP/proofs/3b6aeabae1e044a2927af77a83d131ba.E_input', t123_zfmisc_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/3b6aeabae1e044a2927af77a83d131ba.E_input', t48_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/3b6aeabae1e044a2927af77a83d131ba.E_input', commutativity_k3_xboole_0)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/3b6aeabae1e044a2927af77a83d131ba.E_input', t4_xboole_0)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/3b6aeabae1e044a2927af77a83d131ba.E_input', t36_xboole_1)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/3b6aeabae1e044a2927af77a83d131ba.E_input', t2_boole)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/3b6aeabae1e044a2927af77a83d131ba.E_input', l32_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/3b6aeabae1e044a2927af77a83d131ba.E_input', t3_boole)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/3b6aeabae1e044a2927af77a83d131ba.E_input', t113_zfmisc_1)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/3b6aeabae1e044a2927af77a83d131ba.E_input', t83_xboole_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/3b6aeabae1e044a2927af77a83d131ba.E_input', t7_xboole_0)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/3b6aeabae1e044a2927af77a83d131ba.E_input', d10_xboole_0)).
fof(t32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k4_xboole_0(X2,X1)=>X1=X2), file('data/MPTP/proofs/3b6aeabae1e044a2927af77a83d131ba.E_input', t32_xboole_1)).
fof(c_0_16, negated_conjecture, ~(![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4))))), inference(assume_negation,[status(cth)],[t134_zfmisc_1])).
fof(c_0_17, plain, ![X36, X37, X38]:((~r1_tarski(k2_zfmisc_1(X37,X36),k2_zfmisc_1(X38,X36))|r1_tarski(X37,X38)|X36=k1_xboole_0)&(~r1_tarski(k2_zfmisc_1(X36,X37),k2_zfmisc_1(X36,X38))|r1_tarski(X37,X38)|X36=k1_xboole_0)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t117_zfmisc_1])])])])).
fof(c_0_18, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k2_zfmisc_1(esk3_0,esk4_0)&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk1_0!=esk3_0|esk2_0!=esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).
cnf(c_0_19, plain, (r1_tarski(X2,X3)|X1=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k2_zfmisc_1(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_21, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_22, plain, ![X28, X29, X30]:(k2_zfmisc_1(k4_xboole_0(X28,X29),X30)=k4_xboole_0(k2_zfmisc_1(X28,X30),k2_zfmisc_1(X29,X30))&k2_zfmisc_1(X30,k4_xboole_0(X28,X29))=k4_xboole_0(k2_zfmisc_1(X30,X28),k2_zfmisc_1(X30,X29))), inference(variable_rename,[status(thm)],[t125_zfmisc_1])).
fof(c_0_23, plain, ![X41, X42, X43, X44]:k2_zfmisc_1(k3_xboole_0(X41,X42),k3_xboole_0(X43,X44))=k3_xboole_0(k2_zfmisc_1(X41,X43),k2_zfmisc_1(X42,X44)), inference(variable_rename,[status(thm)],[t123_zfmisc_1])).
fof(c_0_24, plain, ![X31, X32]:k4_xboole_0(X31,k4_xboole_0(X31,X32))=k3_xboole_0(X31,X32), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_25, negated_conjecture, (r1_tarski(esk2_0,X1)|~r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])).
cnf(c_0_26, plain, (k2_zfmisc_1(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_29, plain, ![X26, X27]:k3_xboole_0(X26,X27)=k3_xboole_0(X27,X26), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_30, plain, ![X11, X12, X14, X15, X16]:((r1_xboole_0(X11,X12)|r2_hidden(esk5_2(X11,X12),k3_xboole_0(X11,X12)))&(~r2_hidden(X16,k3_xboole_0(X14,X15))|~r1_xboole_0(X14,X15))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
cnf(c_0_31, negated_conjecture, (r1_tarski(esk2_0,k4_xboole_0(X1,X2))|~r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(esk1_0,X1),k2_zfmisc_1(esk1_0,X2)))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_32, plain, (k4_xboole_0(k2_zfmisc_1(X1,X3),k4_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)))=k2_zfmisc_1(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k4_xboole_0(X3,k4_xboole_0(X3,X4)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_28]), c_0_28]), c_0_28])).
cnf(c_0_33, plain, (k2_zfmisc_1(k4_xboole_0(X1,X2),X3)=k4_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_34, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_35, plain, ![X39, X40]:r1_tarski(k4_xboole_0(X39,X40),X39), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_36, plain, ![X25]:k3_xboole_0(X25,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
cnf(c_0_37, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_38, plain, ![X9, X10]:((k4_xboole_0(X9,X10)!=k1_xboole_0|r1_tarski(X9,X10))&(~r1_tarski(X9,X10)|k4_xboole_0(X9,X10)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_39, negated_conjecture, (r1_tarski(esk2_0,k4_xboole_0(esk2_0,X1))|~r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,X1)))), inference(spm,[status(thm)],[c_0_31, c_0_20])).
cnf(c_0_40, plain, (k4_xboole_0(k4_xboole_0(k2_zfmisc_1(X1,X2),k4_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))),k4_xboole_0(k4_xboole_0(k2_zfmisc_1(X1,X2),k4_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))),k4_xboole_0(k2_zfmisc_1(X4,X2),k4_xboole_0(k2_zfmisc_1(X4,X2),k2_zfmisc_1(X4,X3)))))=k4_xboole_0(k2_zfmisc_1(X1,X2),k4_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X4,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_33]), c_0_26]), c_0_26]), c_0_33]), c_0_26]), c_0_26]), c_0_26]), c_0_26])).
cnf(c_0_41, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_28]), c_0_28])).
cnf(c_0_42, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_43, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_44, plain, ![X35]:k4_xboole_0(X35,k1_xboole_0)=X35, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_45, plain, ![X33, X34]:((k2_zfmisc_1(X33,X34)!=k1_xboole_0|(X33=k1_xboole_0|X34=k1_xboole_0))&((X33!=k1_xboole_0|k2_zfmisc_1(X33,X34)=k1_xboole_0)&(X34!=k1_xboole_0|k2_zfmisc_1(X33,X34)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_46, plain, (~r1_xboole_0(X2,X3)|~r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))), inference(rw,[status(thm)],[c_0_37, c_0_28])).
cnf(c_0_47, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk5_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_48, plain, ![X17, X18]:((~r1_xboole_0(X17,X18)|k4_xboole_0(X17,X18)=X17)&(k4_xboole_0(X17,X18)!=X17|r1_xboole_0(X17,X18))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
cnf(c_0_49, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, negated_conjecture, (r1_tarski(esk2_0,k4_xboole_0(esk2_0,k4_xboole_0(X1,X2)))|~r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(esk1_0,X1),k2_zfmisc_1(esk1_0,X2))))), inference(spm,[status(thm)],[c_0_39, c_0_26])).
cnf(c_0_51, plain, (k4_xboole_0(k4_xboole_0(k2_zfmisc_1(X1,X2),k4_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))),k4_xboole_0(k4_xboole_0(k2_zfmisc_1(X1,X2),k4_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))),k4_xboole_0(k2_zfmisc_1(X4,X3),k4_xboole_0(k2_zfmisc_1(X4,X3),k2_zfmisc_1(X4,X2)))))=k4_xboole_0(k2_zfmisc_1(X1,X2),k4_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X4,X3)))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_52, plain, (r1_tarski(k4_xboole_0(k2_zfmisc_1(X1,X2),k4_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))),k4_xboole_0(k2_zfmisc_1(X1,X2),k4_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X4))))), inference(spm,[status(thm)],[c_0_42, c_0_40])).
cnf(c_0_53, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k1_xboole_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_43, c_0_28])).
cnf(c_0_54, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_55, plain, (X1=k1_xboole_0|X2=k1_xboole_0|k2_zfmisc_1(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_56, plain, (~r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_46, c_0_41])).
cnf(c_0_57, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk5_2(X1,X2),k4_xboole_0(X1,k4_xboole_0(X1,X2)))), inference(rw,[status(thm)],[c_0_47, c_0_28])).
cnf(c_0_58, plain, (r1_xboole_0(X1,X2)|k4_xboole_0(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_49, c_0_42])).
fof(c_0_60, plain, ![X19]:(X19=k1_xboole_0|r2_hidden(esk6_1(X19),X19)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_61, negated_conjecture, (r1_tarski(esk2_0,k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,X1)))|~r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,X1))))), inference(spm,[status(thm)],[c_0_50, c_0_20])).
cnf(c_0_62, plain, (k4_xboole_0(k2_zfmisc_1(X1,X2),k4_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4)))=k4_xboole_0(k2_zfmisc_1(X1,X4),k4_xboole_0(k2_zfmisc_1(X1,X4),k2_zfmisc_1(X3,X2)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_51])).
cnf(c_0_63, negated_conjecture, (r1_tarski(k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(X1,X2))),k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,X2))))), inference(spm,[status(thm)],[c_0_52, c_0_20])).
cnf(c_0_64, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_65, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|X3=k1_xboole_0|k4_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_55, c_0_26])).
cnf(c_0_66, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_67, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_58, c_0_54])).
cnf(c_0_68, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r1_xboole_0(X2,k4_xboole_0(X2,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_59]), c_0_54])).
cnf(c_0_69, plain, (X1=k1_xboole_0|r2_hidden(esk6_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_70, negated_conjecture, (r1_tarski(esk2_0,k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,X1)))|~r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(esk3_0,X1),k4_xboole_0(k2_zfmisc_1(esk3_0,X1),k2_zfmisc_1(esk1_0,esk4_0))))), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_71, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,esk4_0))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_54])).
cnf(c_0_72, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X1,X2)))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_59]), c_0_54])).
cnf(c_0_73, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|k4_xboole_0(X3,X4)=k1_xboole_0|k4_xboole_0(k4_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3)),k4_xboole_0(k2_zfmisc_1(X1,X4),k2_zfmisc_1(X2,X4)))!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_33]), c_0_33])).
cnf(c_0_74, plain, (k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_75, plain, (r1_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_76, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_77, negated_conjecture, (r1_tarski(esk2_0,k4_xboole_0(esk4_0,k4_xboole_0(esk4_0,esk2_0)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_41])).
cnf(c_0_78, plain, (k4_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X1)))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_72, c_0_41])).
cnf(c_0_79, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|X3=k1_xboole_0|k4_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_55, c_0_33])).
cnf(c_0_80, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_81, negated_conjecture, (k4_xboole_0(X1,esk2_0)=k1_xboole_0|k4_xboole_0(X2,esk1_0)=k1_xboole_0|k4_xboole_0(k4_xboole_0(k2_zfmisc_1(X2,X1),k2_zfmisc_1(esk1_0,X1)),k4_xboole_0(k2_zfmisc_1(X2,esk2_0),k2_zfmisc_1(esk3_0,esk4_0)))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_73, c_0_20])).
cnf(c_0_82, negated_conjecture, (k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,esk4_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_71]), c_0_72])).
cnf(c_0_83, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_84, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_xboole_0|~r1_xboole_0(X2,k4_xboole_0(X1,k4_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_76, c_0_41])).
cnf(c_0_85, negated_conjecture, (k4_xboole_0(esk2_0,esk4_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_77]), c_0_78])).
cnf(c_0_86, negated_conjecture, (k4_xboole_0(esk1_0,X1)=k1_xboole_0|k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(X1,esk2_0))!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_20]), c_0_80])).
cnf(c_0_87, negated_conjecture, (k4_xboole_0(esk3_0,esk1_0)=k1_xboole_0|k4_xboole_0(esk4_0,esk2_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_83])])).
cnf(c_0_88, negated_conjecture, (~r1_xboole_0(esk4_0,esk2_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_54]), c_0_54]), c_0_80])).
cnf(c_0_89, negated_conjecture, (k4_xboole_0(esk3_0,esk1_0)=k1_xboole_0|esk4_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_79, c_0_82])).
cnf(c_0_90, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_91, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_92, negated_conjecture, (k4_xboole_0(esk1_0,k4_xboole_0(X1,X2))=k1_xboole_0|k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(X1,esk2_0),k2_zfmisc_1(X2,esk2_0)))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_86, c_0_33])).
cnf(c_0_93, plain, (r2_hidden(esk5_2(k4_xboole_0(k2_zfmisc_1(X1,X2),k4_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))),k4_xboole_0(k4_xboole_0(k2_zfmisc_1(X1,X2),k4_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))),k4_xboole_0(k2_zfmisc_1(X4,X2),k4_xboole_0(k2_zfmisc_1(X4,X2),k2_zfmisc_1(X4,X3))))),k4_xboole_0(k4_xboole_0(k2_zfmisc_1(X1,X2),k4_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))),k4_xboole_0(k2_zfmisc_1(X1,X2),k4_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X4,X3)))))|r1_xboole_0(k4_xboole_0(k2_zfmisc_1(X1,X2),k4_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))),k4_xboole_0(k4_xboole_0(k2_zfmisc_1(X1,X2),k4_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))),k4_xboole_0(k2_zfmisc_1(X4,X2),k4_xboole_0(k2_zfmisc_1(X4,X2),k2_zfmisc_1(X4,X3)))))), inference(spm,[status(thm)],[c_0_57, c_0_40])).
cnf(c_0_94, negated_conjecture, (k4_xboole_0(esk3_0,esk1_0)=k1_xboole_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_87]), c_0_88]), c_0_89])).
cnf(c_0_95, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_90])).
cnf(c_0_96, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_91])).
cnf(c_0_97, plain, (~r2_hidden(X1,k1_xboole_0)|~r1_xboole_0(X2,k1_xboole_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_54]), c_0_64])).
cnf(c_0_98, negated_conjecture, (k4_xboole_0(esk1_0,k4_xboole_0(X1,k4_xboole_0(X2,X3)))=k1_xboole_0|k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(X1,esk2_0),k4_xboole_0(k2_zfmisc_1(X2,esk2_0),k2_zfmisc_1(X3,esk2_0))))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_92, c_0_33])).
cnf(c_0_99, negated_conjecture, (r2_hidden(esk5_2(k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,X1))),k4_xboole_0(k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,X1))),k4_xboole_0(k2_zfmisc_1(X2,esk2_0),k4_xboole_0(k2_zfmisc_1(X2,esk2_0),k2_zfmisc_1(X2,X1))))),k4_xboole_0(k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,X1))),k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(X2,X1)))))|r1_xboole_0(k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,X1))),k4_xboole_0(k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,X1))),k4_xboole_0(k2_zfmisc_1(X2,esk2_0),k4_xboole_0(k2_zfmisc_1(X2,esk2_0),k2_zfmisc_1(X2,X1)))))), inference(spm,[status(thm)],[c_0_93, c_0_20])).
cnf(c_0_100, negated_conjecture, (k4_xboole_0(k2_zfmisc_1(esk3_0,X1),k2_zfmisc_1(esk1_0,X1))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_94]), c_0_95])).
cnf(c_0_101, negated_conjecture, (k4_xboole_0(k2_zfmisc_1(X1,esk2_0),k2_zfmisc_1(X1,esk4_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_85]), c_0_96])).
cnf(c_0_102, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_67])])).
cnf(c_0_103, negated_conjecture, (k4_xboole_0(esk1_0,k4_xboole_0(X1,k4_xboole_0(X2,esk1_0)))=k1_xboole_0|k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(X1,esk2_0),k4_xboole_0(k2_zfmisc_1(X2,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_98, c_0_20])).
cnf(c_0_104, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_46, c_0_69])).
cnf(c_0_105, negated_conjecture, (r1_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk3_0,esk2_0)))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_64]), c_0_100]), c_0_54]), c_0_100]), c_0_54]), c_0_101]), c_0_54]), c_0_100]), c_0_54]), c_0_54]), c_0_64]), c_0_100]), c_0_54]), c_0_100]), c_0_54]), c_0_101]), c_0_54]), c_0_102])).
fof(c_0_106, plain, ![X21, X22]:(((r1_tarski(X21,X22)|X21!=X22)&(r1_tarski(X22,X21)|X21!=X22))&(~r1_tarski(X21,X22)|~r1_tarski(X22,X21)|X21=X22)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_107, plain, ![X23, X24]:(k4_xboole_0(X23,X24)!=k4_xboole_0(X24,X23)|X23=X24), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t32_xboole_1])])).
cnf(c_0_108, negated_conjecture, (k4_xboole_0(esk1_0,X1)=k1_xboole_0|k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k4_xboole_0(k2_zfmisc_1(X1,esk4_0),k4_xboole_0(k2_zfmisc_1(X1,esk4_0),k2_zfmisc_1(esk3_0,esk2_0))))!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_62]), c_0_78])).
cnf(c_0_109, negated_conjecture, (k4_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk3_0,esk2_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_72])).
cnf(c_0_110, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_111, plain, (X1=X2|k4_xboole_0(X1,X2)!=k4_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_112, negated_conjecture, (k4_xboole_0(esk1_0,esk3_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_54]), c_0_64])])).
cnf(c_0_113, plain, (k4_xboole_0(X1,X2)=X1|~r1_tarski(X1,k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_110, c_0_42])).
cnf(c_0_114, negated_conjecture, (r1_tarski(X1,esk2_0)|~r1_tarski(k2_zfmisc_1(esk1_0,X1),k2_zfmisc_1(esk3_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])).
cnf(c_0_115, negated_conjecture, (esk1_0=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_94])])).
cnf(c_0_116, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_117, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=X2|~r1_tarski(X2,k4_xboole_0(X1,k4_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_113, c_0_41])).
cnf(c_0_118, negated_conjecture, (r1_tarski(X1,esk2_0)|~r1_tarski(k2_zfmisc_1(esk3_0,X1),k2_zfmisc_1(esk3_0,esk4_0))), inference(rw,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_119, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_116])).
cnf(c_0_120, negated_conjecture, (esk1_0!=esk3_0|esk2_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_121, negated_conjecture, (esk2_0=esk4_0|~r1_tarski(esk4_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_85]), c_0_54]), c_0_54])).
cnf(c_0_122, negated_conjecture, (r1_tarski(esk4_0,esk2_0)), inference(spm,[status(thm)],[c_0_118, c_0_119])).
cnf(c_0_123, negated_conjecture, (esk2_0!=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_120, c_0_115])])).
cnf(c_0_124, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_121, c_0_122])]), c_0_123]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 125
# Proof object clause steps            : 92
# Proof object formula steps           : 33
# Proof object conjectures             : 41
# Proof object clause conjectures      : 38
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 53
# Proof object simplifying inferences  : 75
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 27
# Processed clauses                    : 4588
# ...of these trivial                  : 149
# ...subsumed                          : 3532
# ...remaining for further processing  : 907
# Other redundant clauses eliminated   : 19
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 22
# Backward-rewritten                   : 405
# Generated clauses                    : 40305
# ...of the previous two non-trivial   : 31254
# Contextual simplify-reflections      : 1
# Paramodulations                      : 40274
# Factorizations                       : 0
# Equation resolutions                 : 31
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 476
#    Positive orientable unit clauses  : 76
#    Positive unorientable unit clauses: 14
#    Negative unit clauses             : 13
#    Non-unit-clauses                  : 373
# Current number of unprocessed clauses: 26161
# ...number of literals in the above   : 48952
# Current number of archived formulas  : 0
# Current number of archived clauses   : 428
# Clause-clause subsumption calls (NU) : 62567
# Rec. Clause-clause subsumption calls : 61906
# Non-unit clause-clause subsumptions  : 2075
# Unit Clause-clause subsumption calls : 3204
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1125
# BW rewrite match successes           : 227
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1570427

# -------------------------------------------------
# User time                : 0.405 s
# System time              : 0.013 s
# Total time               : 0.418 s
# Maximum resident set size: 28596 pages
