# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2)))))))), file('data/MPTP/proofs/3c727a9bd31d4f88b102f7ab9f5912eb.E_input', t7_yellow_0)).
fof(t15_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4))))))), file('data/MPTP/proofs/3c727a9bd31d4f88b102f7ab9f5912eb.E_input', t15_yellow_0)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/3c727a9bd31d4f88b102f7ab9f5912eb.E_input', d1_tarski)).
fof(t38_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))&r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))))), file('data/MPTP/proofs/3c727a9bd31d4f88b102f7ab9f5912eb.E_input', t38_yellow_0)).
fof(t16_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3))))))), file('data/MPTP/proofs/3c727a9bd31d4f88b102f7ab9f5912eb.E_input', t16_yellow_0)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/3c727a9bd31d4f88b102f7ab9f5912eb.E_input', d8_lattice3)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/3c727a9bd31d4f88b102f7ab9f5912eb.E_input', t24_orders_2)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/3c727a9bd31d4f88b102f7ab9f5912eb.E_input', redefinition_k6_domain_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/3c727a9bd31d4f88b102f7ab9f5912eb.E_input', d1_xboole_0)).
fof(d9_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)<=>r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1)))))), file('data/MPTP/proofs/3c727a9bd31d4f88b102f7ab9f5912eb.E_input', d9_orders_2)).
fof(cc3_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/3c727a9bd31d4f88b102f7ab9f5912eb.E_input', cc3_relset_1)).
fof(dt_u1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))), file('data/MPTP/proofs/3c727a9bd31d4f88b102f7ab9f5912eb.E_input', dt_u1_orders_2)).
fof(c_0_12, plain, ![X5, X6, X7]:((~r1_lattice3(X5,k1_tarski(X7),X6)|r1_orders_2(X5,X6,X7)|~m1_subset_1(X7,u1_struct_0(X5))|~m1_subset_1(X6,u1_struct_0(X5))|~l1_orders_2(X5))&((~r1_orders_2(X5,X6,X7)|r1_lattice3(X5,k1_tarski(X7),X6)|~m1_subset_1(X7,u1_struct_0(X5))|~m1_subset_1(X6,u1_struct_0(X5))|~l1_orders_2(X5))&((~r2_lattice3(X5,k1_tarski(X7),X6)|r1_orders_2(X5,X7,X6)|~m1_subset_1(X7,u1_struct_0(X5))|~m1_subset_1(X6,u1_struct_0(X5))|~l1_orders_2(X5))&(~r1_orders_2(X5,X7,X6)|r2_lattice3(X5,k1_tarski(X7),X6)|~m1_subset_1(X7,u1_struct_0(X5))|~m1_subset_1(X6,u1_struct_0(X5))|~l1_orders_2(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])).
fof(c_0_13, plain, ![X8, X9, X11, X12, X13]:(((m1_subset_1(esk1_2(X8,X9),u1_struct_0(X8))|~r1_yellow_0(X8,X9)|(~v5_orders_2(X8)|~l1_orders_2(X8)))&((r2_lattice3(X8,X9,esk1_2(X8,X9))|~r1_yellow_0(X8,X9)|(~v5_orders_2(X8)|~l1_orders_2(X8)))&(~m1_subset_1(X11,u1_struct_0(X8))|(~r2_lattice3(X8,X9,X11)|r1_orders_2(X8,esk1_2(X8,X9),X11))|~r1_yellow_0(X8,X9)|(~v5_orders_2(X8)|~l1_orders_2(X8)))))&((m1_subset_1(esk2_3(X8,X12,X13),u1_struct_0(X8))|~r2_lattice3(X8,X12,X13)|~m1_subset_1(X13,u1_struct_0(X8))|r1_yellow_0(X8,X12)|(~v5_orders_2(X8)|~l1_orders_2(X8)))&((r2_lattice3(X8,X12,esk2_3(X8,X12,X13))|~r2_lattice3(X8,X12,X13)|~m1_subset_1(X13,u1_struct_0(X8))|r1_yellow_0(X8,X12)|(~v5_orders_2(X8)|~l1_orders_2(X8)))&(~r1_orders_2(X8,X13,esk2_3(X8,X12,X13))|~r2_lattice3(X8,X12,X13)|~m1_subset_1(X13,u1_struct_0(X8))|r1_yellow_0(X8,X12)|(~v5_orders_2(X8)|~l1_orders_2(X8)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_yellow_0])])])])])])).
fof(c_0_14, plain, ![X15, X16, X17, X18, X19, X20]:(((~r2_hidden(X17,X16)|X17=X15|X16!=k1_tarski(X15))&(X18!=X15|r2_hidden(X18,X16)|X16!=k1_tarski(X15)))&((~r2_hidden(esk3_2(X19,X20),X20)|esk3_2(X19,X20)!=X19|X20=k1_tarski(X19))&(r2_hidden(esk3_2(X19,X20),X20)|esk3_2(X19,X20)=X19|X20=k1_tarski(X19)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_15, plain, (r1_orders_2(X1,X2,X3)|~r2_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, plain, (r2_lattice3(X1,X2,esk2_3(X1,X2,X3))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, plain, (m1_subset_1(esk2_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_18, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))&r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2)))))), inference(assume_negation,[status(cth)],[t38_yellow_0])).
fof(c_0_19, plain, ![X24, X25, X27, X28, X29]:(((m1_subset_1(esk6_2(X24,X25),u1_struct_0(X24))|~r2_yellow_0(X24,X25)|(~v5_orders_2(X24)|~l1_orders_2(X24)))&((r1_lattice3(X24,X25,esk6_2(X24,X25))|~r2_yellow_0(X24,X25)|(~v5_orders_2(X24)|~l1_orders_2(X24)))&(~m1_subset_1(X27,u1_struct_0(X24))|(~r1_lattice3(X24,X25,X27)|r1_orders_2(X24,X27,esk6_2(X24,X25)))|~r2_yellow_0(X24,X25)|(~v5_orders_2(X24)|~l1_orders_2(X24)))))&((m1_subset_1(esk7_3(X24,X28,X29),u1_struct_0(X24))|~r1_lattice3(X24,X28,X29)|~m1_subset_1(X29,u1_struct_0(X24))|r2_yellow_0(X24,X28)|(~v5_orders_2(X24)|~l1_orders_2(X24)))&((r1_lattice3(X24,X28,esk7_3(X24,X28,X29))|~r1_lattice3(X24,X28,X29)|~m1_subset_1(X29,u1_struct_0(X24))|r2_yellow_0(X24,X28)|(~v5_orders_2(X24)|~l1_orders_2(X24)))&(~r1_orders_2(X24,esk7_3(X24,X28,X29),X29)|~r1_lattice3(X24,X28,X29)|~m1_subset_1(X29,u1_struct_0(X24))|r2_yellow_0(X24,X28)|(~v5_orders_2(X24)|~l1_orders_2(X24)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_yellow_0])])])])])])).
cnf(c_0_20, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_21, plain, ![X31, X32, X33, X34]:((~r1_lattice3(X31,X32,X33)|(~m1_subset_1(X34,u1_struct_0(X31))|(~r2_hidden(X34,X32)|r1_orders_2(X31,X33,X34)))|~m1_subset_1(X33,u1_struct_0(X31))|~l1_orders_2(X31))&((m1_subset_1(esk8_3(X31,X32,X33),u1_struct_0(X31))|r1_lattice3(X31,X32,X33)|~m1_subset_1(X33,u1_struct_0(X31))|~l1_orders_2(X31))&((r2_hidden(esk8_3(X31,X32,X33),X32)|r1_lattice3(X31,X32,X33)|~m1_subset_1(X33,u1_struct_0(X31))|~l1_orders_2(X31))&(~r1_orders_2(X31,X33,esk8_3(X31,X32,X33))|r1_lattice3(X31,X32,X33)|~m1_subset_1(X33,u1_struct_0(X31))|~l1_orders_2(X31))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
cnf(c_0_22, plain, (r1_yellow_0(X1,X3)|~r1_orders_2(X1,X2,esk2_3(X1,X3,X2))|~r2_lattice3(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23, plain, (r1_yellow_0(X1,k1_tarski(X2))|r1_orders_2(X1,X2,esk2_3(X1,k1_tarski(X2),X3))|~v5_orders_2(X1)|~r2_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])).
fof(c_0_24, plain, ![X36, X37]:(v2_struct_0(X36)|(~v3_orders_2(X36)|~l1_orders_2(X36))|(~m1_subset_1(X37,u1_struct_0(X36))|r1_orders_2(X36,X37,X37))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_25, negated_conjecture, ((~v2_struct_0(esk4_0)&(v3_orders_2(esk4_0)&(v5_orders_2(esk4_0)&l1_orders_2(esk4_0))))&(m1_subset_1(esk5_0,u1_struct_0(esk4_0))&(~r1_yellow_0(esk4_0,k6_domain_1(u1_struct_0(esk4_0),esk5_0))|~r2_yellow_0(esk4_0,k6_domain_1(u1_struct_0(esk4_0),esk5_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_18])])])])).
cnf(c_0_26, plain, (r1_orders_2(X1,X3,X2)|~r1_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_27, plain, (r1_lattice3(X1,X2,esk7_3(X1,X2,X3))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (X1=X2|~r2_hidden(X1,k1_tarski(X2))), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_30, plain, (r2_hidden(esk8_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_31, plain, ![X38, X39]:(v1_xboole_0(X38)|~m1_subset_1(X39,X38)|k6_domain_1(X38,X39)=k1_tarski(X39)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_32, plain, (r1_yellow_0(X1,k1_tarski(X2))|~v5_orders_2(X1)|~r2_lattice3(X1,k1_tarski(X2),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_33, plain, (r2_lattice3(X1,k1_tarski(X2),X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_34, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_35, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_36, negated_conjecture, (v3_orders_2(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_37, negated_conjecture, (l1_orders_2(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_38, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_39, plain, (r2_yellow_0(X1,X2)|~r1_orders_2(X1,esk7_3(X1,X2,X3),X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_40, plain, (r2_yellow_0(X1,k1_tarski(X2))|r1_orders_2(X1,esk7_3(X1,k1_tarski(X2),X3),X2)|~v5_orders_2(X1)|~r1_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])).
cnf(c_0_41, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk8_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_42, plain, (esk8_3(X1,k1_tarski(X2),X3)=X2|r1_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
fof(c_0_43, plain, ![X40, X41, X42]:((~v1_xboole_0(X40)|~r2_hidden(X41,X40))&(r2_hidden(esk9_1(X42),X42)|v1_xboole_0(X42))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_44, plain, ![X44, X45, X46]:((~r1_orders_2(X44,X45,X46)|r2_hidden(k4_tarski(X45,X46),u1_orders_2(X44))|~m1_subset_1(X46,u1_struct_0(X44))|~m1_subset_1(X45,u1_struct_0(X44))|~l1_orders_2(X44))&(~r2_hidden(k4_tarski(X45,X46),u1_orders_2(X44))|r1_orders_2(X44,X45,X46)|~m1_subset_1(X46,u1_struct_0(X44))|~m1_subset_1(X45,u1_struct_0(X44))|~l1_orders_2(X44))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_orders_2])])])])).
fof(c_0_45, plain, ![X47, X48, X49]:(~v1_xboole_0(X47)|(~m1_subset_1(X49,k1_zfmisc_1(k2_zfmisc_1(X47,X48)))|v1_xboole_0(X49))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_relset_1])])])).
fof(c_0_46, plain, ![X50]:(~l1_orders_2(X50)|m1_subset_1(u1_orders_2(X50),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X50),u1_struct_0(X50))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_orders_2])])).
cnf(c_0_47, negated_conjecture, (~r1_yellow_0(esk4_0,k6_domain_1(u1_struct_0(esk4_0),esk5_0))|~r2_yellow_0(esk4_0,k6_domain_1(u1_struct_0(esk4_0),esk5_0))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_48, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_49, plain, (r1_yellow_0(X1,k1_tarski(X2))|~v5_orders_2(X1)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_50, negated_conjecture, (r1_orders_2(esk4_0,esk5_0,esk5_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_51, negated_conjecture, (v5_orders_2(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_52, plain, (r2_yellow_0(X1,k1_tarski(X2))|~v5_orders_2(X1)|~r1_lattice3(X1,k1_tarski(X2),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_53, plain, (r1_lattice3(X1,k1_tarski(X2),X3)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_54, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_55, plain, (r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_56, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_57, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_58, negated_conjecture, (v1_xboole_0(u1_struct_0(esk4_0))|~r2_yellow_0(esk4_0,k1_tarski(esk5_0))|~r1_yellow_0(esk4_0,k1_tarski(esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_35])])).
cnf(c_0_59, negated_conjecture, (r1_yellow_0(esk4_0,k1_tarski(esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51]), c_0_35]), c_0_37])])).
cnf(c_0_60, plain, (r2_yellow_0(X1,k1_tarski(X2))|~v5_orders_2(X1)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_61, plain, (~v1_xboole_0(u1_orders_2(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_62, plain, (v1_xboole_0(u1_orders_2(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_63, negated_conjecture, (v1_xboole_0(u1_struct_0(esk4_0))|~r2_yellow_0(esk4_0,k1_tarski(esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_59])])).
cnf(c_0_64, negated_conjecture, (r2_yellow_0(esk4_0,k1_tarski(esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_50]), c_0_51]), c_0_35]), c_0_37])])).
cnf(c_0_65, plain, (~v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_66, negated_conjecture, (v1_xboole_0(u1_struct_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_64])])).
cnf(c_0_67, negated_conjecture, (~r1_orders_2(esk4_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk4_0))|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_37])])).
cnf(c_0_68, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_50]), c_0_35])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 69
# Proof object clause steps            : 44
# Proof object formula steps           : 25
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 25
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 38
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 38
# Processed clauses                    : 206
# ...of these trivial                  : 0
# ...subsumed                          : 23
# ...remaining for further processing  : 183
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 4
# Generated clauses                    : 373
# ...of the previous two non-trivial   : 332
# Contextual simplify-reflections      : 19
# Paramodulations                      : 362
# Factorizations                       : 4
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 134
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 121
# Current number of unprocessed clauses: 184
# ...number of literals in the above   : 1198
# Current number of archived formulas  : 0
# Current number of archived clauses   : 47
# Clause-clause subsumption calls (NU) : 2885
# Rec. Clause-clause subsumption calls : 346
# Non-unit clause-clause subsumptions  : 44
# Unit Clause-clause subsumption calls : 120
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 15038

# -------------------------------------------------
# User time                : 0.019 s
# System time              : 0.000 s
# Total time               : 0.019 s
# Maximum resident set size: 28700 pages
