# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_filter_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_partfun1(X2,X1)&(v3_relat_2(X2)&(v8_relat_2(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>![X3]:(m1_subset_1(X3,X1)=>![X4]:(m2_filter_1(X4,X1,X2)=>(r3_binop_1(X1,X3,X4)=>r3_binop_1(k8_eqrel_1(X1,X2),k9_eqrel_1(X1,X2,X3),k3_filter_1(X1,X2,X4))))))), file('data/MPTP/proofs/3c97435a4d4742389aa3278e4c05ba0b.E_input', t8_filter_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/3c97435a4d4742389aa3278e4c05ba0b.E_input', cc1_relset_1)).
fof(dt_m2_filter_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&v1_relat_1(X2))=>![X3]:(m2_filter_1(X3,X1,X2)=>(v1_funct_1(X3)&(v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))))), file('data/MPTP/proofs/3c97435a4d4742389aa3278e4c05ba0b.E_input', dt_m2_filter_1)).
fof(d7_binop_1, axiom, ![X1, X2]:(m1_subset_1(X2,X1)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))=>(r3_binop_1(X1,X2,X3)<=>(r1_binop_1(X1,X2,X3)&r2_binop_1(X1,X2,X3))))), file('data/MPTP/proofs/3c97435a4d4742389aa3278e4c05ba0b.E_input', d7_binop_1)).
fof(dt_k9_eqrel_1, axiom, ![X1, X2, X3]:((~(v1_xboole_0(X1))&((v3_relat_2(X2)&(v8_relat_2(X2)&(v1_partfun1(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))&m1_subset_1(X3,X1)))=>m2_subset_1(k9_eqrel_1(X1,X2,X3),k1_zfmisc_1(X1),k8_eqrel_1(X1,X2))), file('data/MPTP/proofs/3c97435a4d4742389aa3278e4c05ba0b.E_input', dt_k9_eqrel_1)).
fof(redefinition_k8_eqrel_1, axiom, ![X1, X2]:((v3_relat_2(X2)&(v8_relat_2(X2)&(v1_partfun1(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>k8_eqrel_1(X1,X2)=k7_eqrel_1(X1,X2)), file('data/MPTP/proofs/3c97435a4d4742389aa3278e4c05ba0b.E_input', redefinition_k8_eqrel_1)).
fof(fc3_eqrel_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&(v3_relat_2(X2)&(v8_relat_2(X2)&(v1_partfun1(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))))=>~(v1_xboole_0(k7_eqrel_1(X1,X2)))), file('data/MPTP/proofs/3c97435a4d4742389aa3278e4c05ba0b.E_input', fc3_eqrel_1)).
fof(t7_filter_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_partfun1(X2,X1)&(v3_relat_2(X2)&(v8_relat_2(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>![X3]:(m1_subset_1(X3,X1)=>![X4]:(m2_filter_1(X4,X1,X2)=>(r2_binop_1(X1,X3,X4)=>r2_binop_1(k8_eqrel_1(X1,X2),k9_eqrel_1(X1,X2,X3),k3_filter_1(X1,X2,X4))))))), file('data/MPTP/proofs/3c97435a4d4742389aa3278e4c05ba0b.E_input', t7_filter_1)).
fof(t6_filter_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_partfun1(X2,X1)&(v3_relat_2(X2)&(v8_relat_2(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>![X3]:(m1_subset_1(X3,X1)=>![X4]:(m2_filter_1(X4,X1,X2)=>(r1_binop_1(X1,X3,X4)=>r1_binop_1(k8_eqrel_1(X1,X2),k9_eqrel_1(X1,X2,X3),k3_filter_1(X1,X2,X4))))))), file('data/MPTP/proofs/3c97435a4d4742389aa3278e4c05ba0b.E_input', t6_filter_1)).
fof(dt_k3_filter_1, axiom, ![X1, X2, X3]:((~(v1_xboole_0(X1))&((v1_partfun1(X2,X1)&(v3_relat_2(X2)&(v8_relat_2(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))&(v1_funct_1(X3)&(v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))))=>(v1_funct_1(k3_filter_1(X1,X2,X3))&(v1_funct_2(k3_filter_1(X1,X2,X3),k2_zfmisc_1(k8_eqrel_1(X1,X2),k8_eqrel_1(X1,X2)),k8_eqrel_1(X1,X2))&m1_subset_1(k3_filter_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k8_eqrel_1(X1,X2),k8_eqrel_1(X1,X2)),k8_eqrel_1(X1,X2))))))), file('data/MPTP/proofs/3c97435a4d4742389aa3278e4c05ba0b.E_input', dt_k3_filter_1)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&(~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(X1))))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('data/MPTP/proofs/3c97435a4d4742389aa3278e4c05ba0b.E_input', redefinition_m2_subset_1)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('data/MPTP/proofs/3c97435a4d4742389aa3278e4c05ba0b.E_input', cc1_subset_1)).
fof(dt_k7_eqrel_1, axiom, ![X1, X2]:((v3_relat_2(X2)&(v8_relat_2(X2)&(v1_partfun1(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>m1_subset_1(k7_eqrel_1(X1,X2),k1_zfmisc_1(k1_zfmisc_1(X1)))), file('data/MPTP/proofs/3c97435a4d4742389aa3278e4c05ba0b.E_input', dt_k7_eqrel_1)).
fof(c_0_13, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_partfun1(X2,X1)&(v3_relat_2(X2)&(v8_relat_2(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>![X3]:(m1_subset_1(X3,X1)=>![X4]:(m2_filter_1(X4,X1,X2)=>(r3_binop_1(X1,X3,X4)=>r3_binop_1(k8_eqrel_1(X1,X2),k9_eqrel_1(X1,X2,X3),k3_filter_1(X1,X2,X4)))))))), inference(assume_negation,[status(cth)],[t8_filter_1])).
fof(c_0_14, plain, ![X9, X10, X11]:(~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|v1_relat_1(X11)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_15, negated_conjecture, (~v1_xboole_0(esk1_0)&((v1_partfun1(esk2_0,esk1_0)&(v3_relat_2(esk2_0)&(v8_relat_2(esk2_0)&m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0))))))&(m1_subset_1(esk3_0,esk1_0)&(m2_filter_1(esk4_0,esk1_0,esk2_0)&(r3_binop_1(esk1_0,esk3_0,esk4_0)&~r3_binop_1(k8_eqrel_1(esk1_0,esk2_0),k9_eqrel_1(esk1_0,esk2_0,esk3_0),k3_filter_1(esk1_0,esk2_0,esk4_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_13])])])])).
fof(c_0_16, plain, ![X12, X13, X14]:((v1_funct_1(X14)|~m2_filter_1(X14,X12,X13)|(v1_xboole_0(X12)|~v1_relat_1(X13)))&((v1_funct_2(X14,k2_zfmisc_1(X12,X12),X12)|~m2_filter_1(X14,X12,X13)|(v1_xboole_0(X12)|~v1_relat_1(X13)))&(m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X12,X12),X12)))|~m2_filter_1(X14,X12,X13)|(v1_xboole_0(X12)|~v1_relat_1(X13))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_filter_1])])])])])).
cnf(c_0_17, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_19, plain, ![X15, X16, X17]:(((r1_binop_1(X15,X16,X17)|~r3_binop_1(X15,X16,X17)|(~v1_funct_1(X17)|(~v1_funct_2(X17,k2_zfmisc_1(X15,X15),X15)|~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X15,X15),X15)))))|~m1_subset_1(X16,X15))&(r2_binop_1(X15,X16,X17)|~r3_binop_1(X15,X16,X17)|(~v1_funct_1(X17)|(~v1_funct_2(X17,k2_zfmisc_1(X15,X15),X15)|~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X15,X15),X15)))))|~m1_subset_1(X16,X15)))&(~r1_binop_1(X15,X16,X17)|~r2_binop_1(X15,X16,X17)|r3_binop_1(X15,X16,X17)|(~v1_funct_1(X17)|(~v1_funct_2(X17,k2_zfmisc_1(X15,X15),X15)|~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X15,X15),X15)))))|~m1_subset_1(X16,X15))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_binop_1])])])])).
cnf(c_0_20, plain, (v1_funct_2(X1,k2_zfmisc_1(X2,X2),X2)|v1_xboole_0(X2)|~m2_filter_1(X1,X2,X3)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, negated_conjecture, (m2_filter_1(esk4_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, negated_conjecture, (v1_relat_1(esk2_0)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_23, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, plain, (v1_funct_1(X1)|v1_xboole_0(X2)|~m2_filter_1(X1,X2,X3)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|v1_xboole_0(X2)|~m2_filter_1(X1,X2,X3)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_26, plain, ![X18, X19, X20]:(v1_xboole_0(X18)|(~v3_relat_2(X19)|(~v8_relat_2(X19)|(~v1_partfun1(X19,X18)|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(X18,X18)))))|~m1_subset_1(X20,X18))|m2_subset_1(k9_eqrel_1(X18,X19,X20),k1_zfmisc_1(X18),k8_eqrel_1(X18,X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k9_eqrel_1])])])).
fof(c_0_27, plain, ![X21, X22]:(~v3_relat_2(X22)|(~v8_relat_2(X22)|(~v1_partfun1(X22,X21)|~m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(X21,X21)))))|k8_eqrel_1(X21,X22)=k7_eqrel_1(X21,X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_eqrel_1])])).
fof(c_0_28, plain, ![X23, X24]:(v1_xboole_0(X23)|(~v3_relat_2(X24)|(~v8_relat_2(X24)|(~v1_partfun1(X24,X23)|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(X23,X23))))))|~v1_xboole_0(k7_eqrel_1(X23,X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc3_eqrel_1])])])).
fof(c_0_29, plain, ![X29, X30, X31, X32]:(v1_xboole_0(X29)|(~v1_partfun1(X30,X29)|(~v3_relat_2(X30)|(~v8_relat_2(X30)|~m1_subset_1(X30,k1_zfmisc_1(k2_zfmisc_1(X29,X29)))))|(~m1_subset_1(X31,X29)|(~m2_filter_1(X32,X29,X30)|(~r2_binop_1(X29,X31,X32)|r2_binop_1(k8_eqrel_1(X29,X30),k9_eqrel_1(X29,X30,X31),k3_filter_1(X29,X30,X32))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_filter_1])])])])).
cnf(c_0_30, plain, (r2_binop_1(X1,X2,X3)|~r3_binop_1(X1,X2,X3)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_31, negated_conjecture, (r3_binop_1(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_32, negated_conjecture, (v1_funct_2(esk4_0,k2_zfmisc_1(esk1_0,esk1_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])]), c_0_23])).
cnf(c_0_33, negated_conjecture, (v1_funct_1(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_21]), c_0_22])]), c_0_23])).
cnf(c_0_34, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0),esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_21]), c_0_22])]), c_0_23])).
cnf(c_0_35, negated_conjecture, (m1_subset_1(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_36, plain, ![X25, X26, X27, X28]:(v1_xboole_0(X25)|(~v1_partfun1(X26,X25)|(~v3_relat_2(X26)|(~v8_relat_2(X26)|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(X25,X25)))))|(~m1_subset_1(X27,X25)|(~m2_filter_1(X28,X25,X26)|(~r1_binop_1(X25,X27,X28)|r1_binop_1(k8_eqrel_1(X25,X26),k9_eqrel_1(X25,X26,X27),k3_filter_1(X25,X26,X28))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_filter_1])])])])).
cnf(c_0_37, plain, (r1_binop_1(X1,X2,X3)|~r3_binop_1(X1,X2,X3)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_38, plain, ![X33, X34, X35]:((v1_funct_1(k3_filter_1(X33,X34,X35))|(v1_xboole_0(X33)|(~v1_partfun1(X34,X33)|(~v3_relat_2(X34)|(~v8_relat_2(X34)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X33,X33)))))|(~v1_funct_1(X35)|(~v1_funct_2(X35,k2_zfmisc_1(X33,X33),X33)|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X33,X33),X33))))))))&((v1_funct_2(k3_filter_1(X33,X34,X35),k2_zfmisc_1(k8_eqrel_1(X33,X34),k8_eqrel_1(X33,X34)),k8_eqrel_1(X33,X34))|(v1_xboole_0(X33)|(~v1_partfun1(X34,X33)|(~v3_relat_2(X34)|(~v8_relat_2(X34)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X33,X33)))))|(~v1_funct_1(X35)|(~v1_funct_2(X35,k2_zfmisc_1(X33,X33),X33)|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X33,X33),X33))))))))&(m1_subset_1(k3_filter_1(X33,X34,X35),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k8_eqrel_1(X33,X34),k8_eqrel_1(X33,X34)),k8_eqrel_1(X33,X34))))|(v1_xboole_0(X33)|(~v1_partfun1(X34,X33)|(~v3_relat_2(X34)|(~v8_relat_2(X34)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X33,X33)))))|(~v1_funct_1(X35)|(~v1_funct_2(X35,k2_zfmisc_1(X33,X33),X33)|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X33,X33),X33)))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_filter_1])])])])).
fof(c_0_39, plain, ![X36, X37, X38]:((~m2_subset_1(X38,X36,X37)|m1_subset_1(X38,X37)|(v1_xboole_0(X36)|(v1_xboole_0(X37)|~m1_subset_1(X37,k1_zfmisc_1(X36)))))&(~m1_subset_1(X38,X37)|m2_subset_1(X38,X36,X37)|(v1_xboole_0(X36)|(v1_xboole_0(X37)|~m1_subset_1(X37,k1_zfmisc_1(X36)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])).
fof(c_0_40, plain, ![X39, X40]:(~v1_xboole_0(X39)|(~m1_subset_1(X40,k1_zfmisc_1(X39))|v1_xboole_0(X40))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
cnf(c_0_41, plain, (v1_xboole_0(X1)|m2_subset_1(k9_eqrel_1(X1,X2,X3),k1_zfmisc_1(X1),k8_eqrel_1(X1,X2))|~v3_relat_2(X2)|~v8_relat_2(X2)|~v1_partfun1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_42, negated_conjecture, (v1_partfun1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_43, negated_conjecture, (v8_relat_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_44, negated_conjecture, (v3_relat_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_45, plain, ![X41, X42]:(~v3_relat_2(X42)|(~v8_relat_2(X42)|(~v1_partfun1(X42,X41)|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(X41,X41)))))|m1_subset_1(k7_eqrel_1(X41,X42),k1_zfmisc_1(k1_zfmisc_1(X41)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_eqrel_1])])).
cnf(c_0_46, plain, (k8_eqrel_1(X2,X1)=k7_eqrel_1(X2,X1)|~v3_relat_2(X1)|~v8_relat_2(X1)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_47, plain, (v1_xboole_0(X1)|~v3_relat_2(X2)|~v8_relat_2(X2)|~v1_partfun1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_xboole_0(k7_eqrel_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_48, plain, (v1_xboole_0(X1)|r2_binop_1(k8_eqrel_1(X1,X2),k9_eqrel_1(X1,X2,X3),k3_filter_1(X1,X2,X4))|~v1_partfun1(X2,X1)|~v3_relat_2(X2)|~v8_relat_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X3,X1)|~m2_filter_1(X4,X1,X2)|~r2_binop_1(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_49, negated_conjecture, (r2_binop_1(esk1_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35])])).
cnf(c_0_50, plain, (v1_xboole_0(X1)|r1_binop_1(k8_eqrel_1(X1,X2),k9_eqrel_1(X1,X2,X3),k3_filter_1(X1,X2,X4))|~v1_partfun1(X2,X1)|~v3_relat_2(X2)|~v8_relat_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X3,X1)|~m2_filter_1(X4,X1,X2)|~r1_binop_1(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_51, negated_conjecture, (r1_binop_1(esk1_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35])])).
cnf(c_0_52, plain, (v1_funct_1(k3_filter_1(X1,X2,X3))|v1_xboole_0(X1)|~v1_partfun1(X2,X1)|~v3_relat_2(X2)|~v8_relat_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_53, plain, (m1_subset_1(X1,X3)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_54, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_55, negated_conjecture, (m2_subset_1(k9_eqrel_1(esk1_0,esk2_0,X1),k1_zfmisc_1(esk1_0),k8_eqrel_1(esk1_0,esk2_0))|~m1_subset_1(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_18]), c_0_43]), c_0_44])]), c_0_23])).
cnf(c_0_56, plain, (m1_subset_1(k7_eqrel_1(X2,X1),k1_zfmisc_1(k1_zfmisc_1(X2)))|~v3_relat_2(X1)|~v8_relat_2(X1)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_57, negated_conjecture, (k7_eqrel_1(esk1_0,esk2_0)=k8_eqrel_1(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_42]), c_0_18]), c_0_43]), c_0_44])])).
cnf(c_0_58, negated_conjecture, (~v1_xboole_0(k7_eqrel_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_42]), c_0_18]), c_0_43]), c_0_44])]), c_0_23])).
cnf(c_0_59, negated_conjecture, (r2_binop_1(k8_eqrel_1(esk1_0,X1),k9_eqrel_1(esk1_0,X1,esk3_0),k3_filter_1(esk1_0,X1,esk4_0))|~m2_filter_1(esk4_0,esk1_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))|~v8_relat_2(X1)|~v3_relat_2(X1)|~v1_partfun1(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_35])]), c_0_23])).
cnf(c_0_60, negated_conjecture, (r1_binop_1(k8_eqrel_1(esk1_0,X1),k9_eqrel_1(esk1_0,X1,esk3_0),k3_filter_1(esk1_0,X1,esk4_0))|~m2_filter_1(esk4_0,esk1_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))|~v8_relat_2(X1)|~v3_relat_2(X1)|~v1_partfun1(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_35])]), c_0_23])).
cnf(c_0_61, negated_conjecture, (v1_funct_1(k3_filter_1(esk1_0,X1,esk4_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))|~v8_relat_2(X1)|~v3_relat_2(X1)|~v1_partfun1(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_32]), c_0_33]), c_0_34])]), c_0_23])).
cnf(c_0_62, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~m2_subset_1(X1,X3,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(csr,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_63, negated_conjecture, (m2_subset_1(k9_eqrel_1(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk1_0),k8_eqrel_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_55, c_0_35])).
cnf(c_0_64, negated_conjecture, (m1_subset_1(k8_eqrel_1(esk1_0,esk2_0),k1_zfmisc_1(k1_zfmisc_1(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_42]), c_0_57]), c_0_18]), c_0_43]), c_0_44])])).
cnf(c_0_65, negated_conjecture, (~v1_xboole_0(k8_eqrel_1(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_58, c_0_57])).
cnf(c_0_66, plain, (m1_subset_1(k3_filter_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k8_eqrel_1(X1,X2),k8_eqrel_1(X1,X2)),k8_eqrel_1(X1,X2))))|v1_xboole_0(X1)|~v1_partfun1(X2,X1)|~v3_relat_2(X2)|~v8_relat_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_67, plain, (r3_binop_1(X1,X2,X3)|~r1_binop_1(X1,X2,X3)|~r2_binop_1(X1,X2,X3)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_68, negated_conjecture, (r2_binop_1(k8_eqrel_1(esk1_0,esk2_0),k9_eqrel_1(esk1_0,esk2_0,esk3_0),k3_filter_1(esk1_0,esk2_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_21]), c_0_18]), c_0_43]), c_0_44]), c_0_42])])).
cnf(c_0_69, negated_conjecture, (r1_binop_1(k8_eqrel_1(esk1_0,esk2_0),k9_eqrel_1(esk1_0,esk2_0,esk3_0),k3_filter_1(esk1_0,esk2_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_21]), c_0_18]), c_0_43]), c_0_44]), c_0_42])])).
cnf(c_0_70, negated_conjecture, (v1_funct_1(k3_filter_1(esk1_0,esk2_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_42]), c_0_18]), c_0_43]), c_0_44])])).
cnf(c_0_71, negated_conjecture, (m1_subset_1(k9_eqrel_1(esk1_0,esk2_0,esk3_0),k8_eqrel_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64])]), c_0_65])).
cnf(c_0_72, negated_conjecture, (~r3_binop_1(k8_eqrel_1(esk1_0,esk2_0),k9_eqrel_1(esk1_0,esk2_0,esk3_0),k3_filter_1(esk1_0,esk2_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_73, negated_conjecture, (m1_subset_1(k3_filter_1(esk1_0,X1,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k8_eqrel_1(esk1_0,X1),k8_eqrel_1(esk1_0,X1)),k8_eqrel_1(esk1_0,X1))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))|~v8_relat_2(X1)|~v3_relat_2(X1)|~v1_partfun1(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_32]), c_0_33]), c_0_34])]), c_0_23])).
cnf(c_0_74, plain, (v1_funct_2(k3_filter_1(X1,X2,X3),k2_zfmisc_1(k8_eqrel_1(X1,X2),k8_eqrel_1(X1,X2)),k8_eqrel_1(X1,X2))|v1_xboole_0(X1)|~v1_partfun1(X2,X1)|~v3_relat_2(X2)|~v8_relat_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_75, negated_conjecture, (~v1_funct_2(k3_filter_1(esk1_0,esk2_0,esk4_0),k2_zfmisc_1(k8_eqrel_1(esk1_0,esk2_0),k8_eqrel_1(esk1_0,esk2_0)),k8_eqrel_1(esk1_0,esk2_0))|~m1_subset_1(k3_filter_1(esk1_0,esk2_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k8_eqrel_1(esk1_0,esk2_0),k8_eqrel_1(esk1_0,esk2_0)),k8_eqrel_1(esk1_0,esk2_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_69]), c_0_70]), c_0_71])]), c_0_72])).
cnf(c_0_76, negated_conjecture, (m1_subset_1(k3_filter_1(esk1_0,esk2_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k8_eqrel_1(esk1_0,esk2_0),k8_eqrel_1(esk1_0,esk2_0)),k8_eqrel_1(esk1_0,esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_42]), c_0_18]), c_0_43]), c_0_44])])).
cnf(c_0_77, negated_conjecture, (v1_funct_2(k3_filter_1(esk1_0,X1,esk4_0),k2_zfmisc_1(k8_eqrel_1(esk1_0,X1),k8_eqrel_1(esk1_0,X1)),k8_eqrel_1(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))|~v8_relat_2(X1)|~v3_relat_2(X1)|~v1_partfun1(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_32]), c_0_33]), c_0_34])]), c_0_23])).
cnf(c_0_78, negated_conjecture, (~v1_funct_2(k3_filter_1(esk1_0,esk2_0,esk4_0),k2_zfmisc_1(k8_eqrel_1(esk1_0,esk2_0),k8_eqrel_1(esk1_0,esk2_0)),k8_eqrel_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_76])])).
cnf(c_0_79, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_42]), c_0_18]), c_0_43]), c_0_44])]), c_0_78]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 80
# Proof object clause steps            : 53
# Proof object formula steps           : 27
# Proof object conjectures             : 37
# Proof object clause conjectures      : 34
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 23
# Proof object simplifying inferences  : 91
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 28
# Processed clauses                    : 94
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 94
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 39
# ...of the previous two non-trivial   : 38
# Contextual simplify-reflections      : 2
# Paramodulations                      : 39
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 64
#    Positive orientable unit clauses  : 23
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 36
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 30
# Clause-clause subsumption calls (NU) : 924
# Rec. Clause-clause subsumption calls : 83
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 25
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5080

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.000 s
# Total time               : 0.013 s
# Maximum resident set size: 28636 pages
