# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/3cf1923925d243a284e1244b22cedbbb.E_input', t3_subset)).
fof(existence_m1_subset_1, axiom, ![X1]:?[X2]:m1_subset_1(X2,X1), file('data/MPTP/proofs/3cf1923925d243a284e1244b22cedbbb.E_input', existence_m1_subset_1)).
fof(t12_mcart_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(k1_tarski(X2),X3))=>(k1_mcart_1(X1)=X2&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/3cf1923925d243a284e1244b22cedbbb.E_input', t12_mcart_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/3cf1923925d243a284e1244b22cedbbb.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/3cf1923925d243a284e1244b22cedbbb.E_input', t70_enumset1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/3cf1923925d243a284e1244b22cedbbb.E_input', t113_zfmisc_1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/3cf1923925d243a284e1244b22cedbbb.E_input', t7_ordinal1)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/3cf1923925d243a284e1244b22cedbbb.E_input', d2_tarski)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/3cf1923925d243a284e1244b22cedbbb.E_input', d3_tarski)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/3cf1923925d243a284e1244b22cedbbb.E_input', d5_xboole_0)).
fof(t20_zfmisc_1, axiom, ![X1, X2]:(k4_xboole_0(k1_tarski(X1),k1_tarski(X2))=k1_tarski(X1)<=>~(X1=X2)), file('data/MPTP/proofs/3cf1923925d243a284e1244b22cedbbb.E_input', t20_zfmisc_1)).
fof(t20_mcart_1, conjecture, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1)))), file('data/MPTP/proofs/3cf1923925d243a284e1244b22cedbbb.E_input', t20_mcart_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/3cf1923925d243a284e1244b22cedbbb.E_input', d1_tarski)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/3cf1923925d243a284e1244b22cedbbb.E_input', t7_mcart_1)).
fof(t9_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k4_tarski(X3,X4))))))), file('data/MPTP/proofs/3cf1923925d243a284e1244b22cedbbb.E_input', t9_mcart_1)).
fof(c_0_15, plain, ![X18, X19]:((~m1_subset_1(X18,k1_zfmisc_1(X19))|r1_tarski(X18,X19))&(~r1_tarski(X18,X19)|m1_subset_1(X18,k1_zfmisc_1(X19)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_16, plain, ![X50]:m1_subset_1(esk8_1(X50),X50), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[existence_m1_subset_1])])).
fof(c_0_17, plain, ![X8, X9, X10]:((k1_mcart_1(X8)=X9|~r2_hidden(X8,k2_zfmisc_1(k1_tarski(X9),X10)))&(r2_hidden(k2_mcart_1(X8),X10)|~r2_hidden(X8,k2_zfmisc_1(k1_tarski(X9),X10)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_mcart_1])])])).
fof(c_0_18, plain, ![X17]:k2_tarski(X17,X17)=k1_tarski(X17), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_19, plain, ![X31, X32]:k1_enumset1(X31,X31,X32)=k2_tarski(X31,X32), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_20, plain, ![X29, X30]:((k2_zfmisc_1(X29,X30)!=k1_xboole_0|(X29=k1_xboole_0|X30=k1_xboole_0))&((X29!=k1_xboole_0|k2_zfmisc_1(X29,X30)=k1_xboole_0)&(X30!=k1_xboole_0|k2_zfmisc_1(X29,X30)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_21, plain, ![X33, X34]:(~r2_hidden(X33,X34)|~r1_tarski(X34,X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_22, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (m1_subset_1(esk8_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (k1_mcart_1(X1)=X2|~r2_hidden(X1,k2_zfmisc_1(k1_tarski(X2),X3))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, plain, (r1_tarski(esk8_1(k1_zfmisc_1(X1)),X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_30, plain, (k1_mcart_1(X1)=X2|~r2_hidden(X1,k2_zfmisc_1(k1_enumset1(X2,X2,X2),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25]), c_0_26])).
cnf(c_0_31, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_32, plain, (~r2_hidden(X1,esk8_1(k1_zfmisc_1(X1)))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_33, plain, (k1_mcart_1(X1)=X2|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
fof(c_0_34, plain, ![X20, X21, X22, X23, X24, X25, X26, X27]:(((~r2_hidden(X23,X22)|(X23=X20|X23=X21)|X22!=k2_tarski(X20,X21))&((X24!=X20|r2_hidden(X24,X22)|X22!=k2_tarski(X20,X21))&(X24!=X21|r2_hidden(X24,X22)|X22!=k2_tarski(X20,X21))))&(((esk5_3(X25,X26,X27)!=X25|~r2_hidden(esk5_3(X25,X26,X27),X27)|X27=k2_tarski(X25,X26))&(esk5_3(X25,X26,X27)!=X26|~r2_hidden(esk5_3(X25,X26,X27),X27)|X27=k2_tarski(X25,X26)))&(r2_hidden(esk5_3(X25,X26,X27),X27)|(esk5_3(X25,X26,X27)=X25|esk5_3(X25,X26,X27)=X26)|X27=k2_tarski(X25,X26)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_35, plain, (~r2_hidden(X1,k1_mcart_1(X2))|~r2_hidden(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
fof(c_0_36, plain, ![X11, X12, X13, X14, X15]:((~r1_tarski(X11,X12)|(~r2_hidden(X13,X11)|r2_hidden(X13,X12)))&((r2_hidden(esk4_2(X14,X15),X14)|r1_tarski(X14,X15))&(~r2_hidden(esk4_2(X14,X15),X15)|r1_tarski(X14,X15)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_37, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_38, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_35, c_0_33])).
cnf(c_0_39, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_40, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X4,X2)), inference(rw,[status(thm)],[c_0_37, c_0_26])).
cnf(c_0_41, plain, (r1_tarski(k1_xboole_0,X1)|~r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_42, plain, (r2_hidden(X1,k1_enumset1(X2,X2,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_40])])).
cnf(c_0_43, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
fof(c_0_44, plain, ![X39, X40, X41, X42, X43, X44, X45, X46]:((((r2_hidden(X42,X39)|~r2_hidden(X42,X41)|X41!=k4_xboole_0(X39,X40))&(~r2_hidden(X42,X40)|~r2_hidden(X42,X41)|X41!=k4_xboole_0(X39,X40)))&(~r2_hidden(X43,X39)|r2_hidden(X43,X40)|r2_hidden(X43,X41)|X41!=k4_xboole_0(X39,X40)))&((~r2_hidden(esk7_3(X44,X45,X46),X46)|(~r2_hidden(esk7_3(X44,X45,X46),X44)|r2_hidden(esk7_3(X44,X45,X46),X45))|X46=k4_xboole_0(X44,X45))&((r2_hidden(esk7_3(X44,X45,X46),X44)|r2_hidden(esk7_3(X44,X45,X46),X46)|X46=k4_xboole_0(X44,X45))&(~r2_hidden(esk7_3(X44,X45,X46),X45)|r2_hidden(esk7_3(X44,X45,X46),X46)|X46=k4_xboole_0(X44,X45))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_45, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_28, c_0_43])).
cnf(c_0_46, plain, (r2_hidden(esk7_3(X1,X2,X3),X1)|r2_hidden(esk7_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_47, plain, ![X48, X49]:((k4_xboole_0(k1_tarski(X48),k1_tarski(X49))!=k1_tarski(X48)|X48!=X49)&(X48=X49|k4_xboole_0(k1_tarski(X48),k1_tarski(X49))=k1_tarski(X48))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t20_zfmisc_1])])])).
fof(c_0_48, negated_conjecture, ~(![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1))))), inference(assume_negation,[status(cth)],[t20_mcart_1])).
cnf(c_0_49, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|r2_hidden(esk7_3(X1,X2,k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
fof(c_0_50, plain, ![X52, X53, X54, X55, X56, X57]:(((~r2_hidden(X54,X53)|X54=X52|X53!=k1_tarski(X52))&(X55!=X52|r2_hidden(X55,X53)|X53!=k1_tarski(X52)))&((~r2_hidden(esk9_2(X56,X57),X57)|esk9_2(X56,X57)!=X56|X57=k1_tarski(X56))&(r2_hidden(esk9_2(X56,X57),X57)|esk9_2(X56,X57)=X56|X57=k1_tarski(X56)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_51, plain, (k4_xboole_0(k1_tarski(X1),k1_tarski(X2))!=k1_tarski(X1)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_52, plain, ![X59, X60]:(k1_mcart_1(k4_tarski(X59,X60))=X59&k2_mcart_1(k4_tarski(X59,X60))=X60), inference(variable_rename,[status(thm)],[t7_mcart_1])).
fof(c_0_53, negated_conjecture, (esk1_0=k4_tarski(esk2_0,esk3_0)&(esk1_0=k1_mcart_1(esk1_0)|esk1_0=k2_mcart_1(esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_48])])])])).
cnf(c_0_54, plain, (X1=k4_xboole_0(k1_xboole_0,X2)|r2_hidden(esk7_3(k1_xboole_0,X2,X1),X1)|~r2_hidden(X3,X4)), inference(spm,[status(thm)],[c_0_38, c_0_46])).
cnf(c_0_55, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_45, c_0_49])).
cnf(c_0_56, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_57, plain, (X1!=X2|k4_xboole_0(k1_enumset1(X1,X1,X1),k1_enumset1(X2,X2,X2))!=k1_enumset1(X1,X1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_25]), c_0_25]), c_0_25]), c_0_26]), c_0_26]), c_0_26])).
cnf(c_0_58, plain, (r2_hidden(esk7_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk7_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_59, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_60, negated_conjecture, (esk1_0=k4_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_53])).
fof(c_0_61, plain, ![X35, X37, X38]:((r2_hidden(esk6_1(X35),X35)|X35=k1_xboole_0)&((~r2_hidden(X37,X35)|esk6_1(X35)!=k4_tarski(X37,X38)|X35=k1_xboole_0)&(~r2_hidden(X38,X35)|esk6_1(X35)!=k4_tarski(X37,X38)|X35=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_mcart_1])])])])])])).
cnf(c_0_62, plain, (X1=k1_xboole_0|r2_hidden(esk7_3(k1_xboole_0,X2,X1),X1)|~r2_hidden(X3,X4)), inference(rw,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_63, plain, (X1=X3|X2!=k1_enumset1(X3,X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_25]), c_0_26])).
cnf(c_0_64, plain, (k4_xboole_0(k1_enumset1(X1,X1,X1),k1_enumset1(X1,X1,X1))!=k1_enumset1(X1,X1,X1)), inference(er,[status(thm)],[c_0_57])).
cnf(c_0_65, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_49]), c_0_45])).
cnf(c_0_66, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_67, negated_conjecture, (esk1_0=k1_mcart_1(esk1_0)|esk1_0=k2_mcart_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_68, negated_conjecture, (k1_mcart_1(esk1_0)=esk2_0), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_69, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk6_1(X2)!=k4_tarski(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_70, plain, (X1=k1_xboole_0|r2_hidden(esk7_3(k1_xboole_0,X2,X1),X1)), inference(spm,[status(thm)],[c_0_62, c_0_42])).
cnf(c_0_71, plain, (X1=X2|~r2_hidden(X1,k1_enumset1(X2,X2,X2))), inference(er,[status(thm)],[c_0_63])).
cnf(c_0_72, plain, (k1_enumset1(X1,X1,X1)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_73, negated_conjecture, (k2_mcart_1(esk1_0)=esk3_0), inference(spm,[status(thm)],[c_0_66, c_0_60])).
cnf(c_0_74, negated_conjecture, (k2_mcart_1(esk1_0)=esk1_0|esk2_0=esk1_0), inference(rw,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_75, plain, (X1=k1_xboole_0|k4_tarski(X2,esk7_3(k1_xboole_0,X3,X1))!=esk6_1(X1)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_76, plain, (esk7_3(k1_xboole_0,X1,k1_enumset1(X2,X2,X2))=X2), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_70]), c_0_72])).
cnf(c_0_77, plain, (r2_hidden(esk6_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_78, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk6_1(X2)!=k4_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_79, negated_conjecture, (esk2_0=esk1_0|esk3_0=esk1_0), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_80, plain, (k4_tarski(X1,X2)!=esk6_1(k1_enumset1(X2,X2,X2))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_72])).
cnf(c_0_81, plain, (esk6_1(k1_enumset1(X1,X1,X1))=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_77]), c_0_72])).
cnf(c_0_82, plain, (X1=k1_xboole_0|k4_tarski(esk7_3(k1_xboole_0,X2,X1),X3)!=esk6_1(X1)), inference(spm,[status(thm)],[c_0_78, c_0_70])).
cnf(c_0_83, negated_conjecture, (k4_tarski(esk2_0,esk1_0)=esk1_0|esk2_0=esk1_0), inference(spm,[status(thm)],[c_0_60, c_0_79])).
cnf(c_0_84, plain, (k4_tarski(X1,X2)!=X2), inference(rw,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_85, plain, (k4_tarski(X1,X2)!=esk6_1(k1_enumset1(X1,X1,X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_76]), c_0_72])).
cnf(c_0_86, negated_conjecture, (esk2_0=esk1_0), inference(sr,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_87, plain, (k4_tarski(X1,X2)!=X1), inference(rw,[status(thm)],[c_0_85, c_0_81])).
cnf(c_0_88, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_86]), c_0_87]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 89
# Proof object clause steps            : 58
# Proof object formula steps           : 31
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 23
# Proof object simplifying inferences  : 29
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 37
# Processed clauses                    : 199
# ...of these trivial                  : 1
# ...subsumed                          : 68
# ...remaining for further processing  : 130
# Other redundant clauses eliminated   : 32
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 11
# Generated clauses                    : 567
# ...of the previous two non-trivial   : 525
# Contextual simplify-reflections      : 0
# Paramodulations                      : 534
# Factorizations                       : 2
# Equation resolutions                 : 32
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 65
#    Positive orientable unit clauses  : 18
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 40
# Current number of unprocessed clauses: 380
# ...number of literals in the above   : 1394
# Current number of archived formulas  : 0
# Current number of archived clauses   : 56
# Clause-clause subsumption calls (NU) : 628
# Rec. Clause-clause subsumption calls : 483
# Non-unit clause-clause subsumptions  : 10
# Unit Clause-clause subsumption calls : 118
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 24
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7608

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 28632 pages
