# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.006 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_mcart_1, axiom, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>![X2]:(X2=k2_mcart_1(X1)<=>![X3, X4]:(X1=k4_tarski(X3,X4)=>X2=X4))), file('data/MPTP/proofs/3d10684eb375496fad467d3cc289daeb.E_input', d2_mcart_1)).
fof(t7_mcart_1, conjecture, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/3d10684eb375496fad467d3cc289daeb.E_input', t7_mcart_1)).
fof(d1_mcart_1, axiom, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>![X2]:(X2=k1_mcart_1(X1)<=>![X3, X4]:(X1=k4_tarski(X3,X4)=>X2=X3))), file('data/MPTP/proofs/3d10684eb375496fad467d3cc289daeb.E_input', d1_mcart_1)).
fof(c_0_3, plain, ![X16, X17, X18, X19, X20, X21, X22]:((X19!=k2_mcart_1(X16)|(X16!=k4_tarski(X20,X21)|X19=X21)|X16!=k4_tarski(X17,X18))&((X16=k4_tarski(esk5_2(X16,X22),esk6_2(X16,X22))|X22=k2_mcart_1(X16)|X16!=k4_tarski(X17,X18))&(X22!=esk6_2(X16,X22)|X22=k2_mcart_1(X16)|X16!=k4_tarski(X17,X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_mcart_1])])])])])])).
fof(c_0_4, negated_conjecture, ~(![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2)), inference(assume_negation,[status(cth)],[t7_mcart_1])).
cnf(c_0_5, plain, (X1=X4|X1!=k2_mcart_1(X2)|X2!=k4_tarski(X3,X4)|X2!=k4_tarski(X5,X6)), inference(split_conjunct,[status(thm)],[c_0_3])).
fof(c_0_6, plain, ![X7, X8, X9, X10, X11, X12, X13]:((X10!=k1_mcart_1(X7)|(X7!=k4_tarski(X11,X12)|X10=X11)|X7!=k4_tarski(X8,X9))&((X7=k4_tarski(esk3_2(X7,X13),esk4_2(X7,X13))|X13=k1_mcart_1(X7)|X7!=k4_tarski(X8,X9))&(X13!=esk3_2(X7,X13)|X13=k1_mcart_1(X7)|X7!=k4_tarski(X8,X9)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_mcart_1])])])])])])).
fof(c_0_7, negated_conjecture, (k1_mcart_1(k4_tarski(esk1_0,esk2_0))!=esk1_0|k2_mcart_1(k4_tarski(esk1_0,esk2_0))!=esk2_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
cnf(c_0_8, plain, (k2_mcart_1(k4_tarski(X1,X2))=X3|k4_tarski(X1,X2)!=k4_tarski(X4,X3)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_5])])).
cnf(c_0_9, plain, (X1=X3|X1!=k1_mcart_1(X2)|X2!=k4_tarski(X3,X4)|X2!=k4_tarski(X5,X6)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, negated_conjecture, (k1_mcart_1(k4_tarski(esk1_0,esk2_0))!=esk1_0|k2_mcart_1(k4_tarski(esk1_0,esk2_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(er,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (k1_mcart_1(k4_tarski(X1,X2))=X3|k4_tarski(X1,X2)!=k4_tarski(X3,X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_9])])).
cnf(c_0_13, negated_conjecture, (k1_mcart_1(k4_tarski(esk1_0,esk2_0))!=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_11])])).
cnf(c_0_14, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(er,[status(thm)],[c_0_12])).
cnf(c_0_15, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_14])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 16
# Proof object clause steps            : 9
# Proof object formula steps           : 7
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 3
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 2
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 7
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 7
# Processed clauses                    : 23
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 23
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 7
# Generated clauses                    : 22
# ...of the previous two non-trivial   : 27
# Contextual simplify-reflections      : 0
# Paramodulations                      : 14
# Factorizations                       : 0
# Equation resolutions                 : 10
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 3
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 1
# Current number of unprocessed clauses: 15
# ...number of literals in the above   : 41
# Current number of archived formulas  : 0
# Current number of archived clauses   : 14
# Clause-clause subsumption calls (NU) : 6
# Rec. Clause-clause subsumption calls : 6
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 708

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.000 s
# Total time               : 0.007 s
# Maximum resident set size: 28636 pages
