# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t75_mcart_1, conjecture, ![X1, X2, X3, X4, X5, X6]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&(~(X5=k1_xboole_0)&(~(X6=k1_xboole_0)&?[X7]:(m1_subset_1(X7,k3_zfmisc_1(X1,X2,X3))&?[X8]:(m1_subset_1(X8,k3_zfmisc_1(X4,X5,X6))&(X7=X8&~((k5_mcart_1(X1,X2,X3,X7)=k5_mcart_1(X4,X5,X6,X8)&(k6_mcart_1(X1,X2,X3,X7)=k6_mcart_1(X4,X5,X6,X8)&k7_mcart_1(X1,X2,X3,X7)=k7_mcart_1(X4,X5,X6,X8)))))))))))))), file('data/MPTP/proofs/3d1d857218c04ab29ded8c7a90f07e3f.E_input', t75_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/3d1d857218c04ab29ded8c7a90f07e3f.E_input', d3_zfmisc_1)).
fof(t50_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))&(k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))&k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4))))))))), file('data/MPTP/proofs/3d1d857218c04ab29ded8c7a90f07e3f.E_input', t50_mcart_1)).
fof(c_0_3, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&(~(X5=k1_xboole_0)&(~(X6=k1_xboole_0)&?[X7]:(m1_subset_1(X7,k3_zfmisc_1(X1,X2,X3))&?[X8]:(m1_subset_1(X8,k3_zfmisc_1(X4,X5,X6))&(X7=X8&~((k5_mcart_1(X1,X2,X3,X7)=k5_mcart_1(X4,X5,X6,X8)&(k6_mcart_1(X1,X2,X3,X7)=k6_mcart_1(X4,X5,X6,X8)&k7_mcart_1(X1,X2,X3,X7)=k7_mcart_1(X4,X5,X6,X8))))))))))))))), inference(assume_negation,[status(cth)],[t75_mcart_1])).
fof(c_0_4, negated_conjecture, (esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&(esk4_0!=k1_xboole_0&(esk5_0!=k1_xboole_0&(esk6_0!=k1_xboole_0&(m1_subset_1(esk7_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))&(m1_subset_1(esk8_0,k3_zfmisc_1(esk4_0,esk5_0,esk6_0))&(esk7_0=esk8_0&(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0)!=k5_mcart_1(esk4_0,esk5_0,esk6_0,esk8_0)|(k6_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0)!=k6_mcart_1(esk4_0,esk5_0,esk6_0,esk8_0)|k7_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0)!=k7_mcart_1(esk4_0,esk5_0,esk6_0,esk8_0)))))))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_3])])])])).
fof(c_0_5, plain, ![X21, X22, X23]:k3_zfmisc_1(X21,X22,X23)=k2_zfmisc_1(k2_zfmisc_1(X21,X22),X23), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_6, plain, ![X17, X18, X19, X20]:((k5_mcart_1(X17,X18,X19,X20)=k1_mcart_1(k1_mcart_1(X20))|~m1_subset_1(X20,k3_zfmisc_1(X17,X18,X19))|X19=k1_xboole_0|X18=k1_xboole_0|X17=k1_xboole_0)&((k6_mcart_1(X17,X18,X19,X20)=k2_mcart_1(k1_mcart_1(X20))|~m1_subset_1(X20,k3_zfmisc_1(X17,X18,X19))|X19=k1_xboole_0|X18=k1_xboole_0|X17=k1_xboole_0)&(k7_mcart_1(X17,X18,X19,X20)=k2_mcart_1(X20)|~m1_subset_1(X20,k3_zfmisc_1(X17,X18,X19))|X19=k1_xboole_0|X18=k1_xboole_0|X17=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_mcart_1])])])])])).
cnf(c_0_7, negated_conjecture, (m1_subset_1(esk8_0,k3_zfmisc_1(esk4_0,esk5_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, negated_conjecture, (m1_subset_1(esk7_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_11, negated_conjecture, (m1_subset_1(esk8_0,k2_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0),esk6_0))), inference(rw,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_12, negated_conjecture, (esk7_0=esk8_0), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_13, plain, (k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14, negated_conjecture, (k5_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0)!=k5_mcart_1(esk4_0,esk5_0,esk6_0,esk8_0)|k6_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0)!=k6_mcart_1(esk4_0,esk5_0,esk6_0,esk8_0)|k7_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0)!=k7_mcart_1(esk4_0,esk5_0,esk6_0,esk8_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_15, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_9, c_0_8])).
cnf(c_0_16, negated_conjecture, (m1_subset_1(esk7_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(rw,[status(thm)],[c_0_10, c_0_8])).
cnf(c_0_17, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_18, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_19, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_20, negated_conjecture, (m1_subset_1(esk7_0,k2_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0),esk6_0))), inference(rw,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_21, negated_conjecture, (esk6_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_22, negated_conjecture, (esk5_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_23, negated_conjecture, (esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_24, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_13, c_0_8])).
cnf(c_0_25, plain, (k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_26, negated_conjecture, (k5_mcart_1(esk4_0,esk5_0,esk6_0,esk7_0)!=k5_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0)|k6_mcart_1(esk4_0,esk5_0,esk6_0,esk7_0)!=k6_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0)|k7_mcart_1(esk4_0,esk5_0,esk6_0,esk7_0)!=k7_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_12]), c_0_12]), c_0_12])).
cnf(c_0_27, negated_conjecture, (k7_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0)=k2_mcart_1(esk7_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_18]), c_0_19])).
cnf(c_0_28, negated_conjecture, (k7_mcart_1(esk4_0,esk5_0,esk6_0,esk7_0)=k2_mcart_1(esk7_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_20]), c_0_21]), c_0_22]), c_0_23])).
cnf(c_0_29, negated_conjecture, (k1_mcart_1(k1_mcart_1(esk7_0))=k5_mcart_1(esk4_0,esk5_0,esk6_0,esk7_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_20]), c_0_21]), c_0_22]), c_0_23])).
cnf(c_0_30, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_25, c_0_8])).
cnf(c_0_31, negated_conjecture, (k5_mcart_1(esk4_0,esk5_0,esk6_0,esk7_0)!=k5_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0)|k6_mcart_1(esk4_0,esk5_0,esk6_0,esk7_0)!=k6_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_27]), c_0_28])])).
cnf(c_0_32, negated_conjecture, (k5_mcart_1(esk4_0,esk5_0,esk6_0,esk7_0)=k5_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0)), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_29])).
cnf(c_0_33, negated_conjecture, (k2_mcart_1(k1_mcart_1(esk7_0))=k6_mcart_1(esk4_0,esk5_0,esk6_0,esk7_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_20]), c_0_21]), c_0_22]), c_0_23])).
cnf(c_0_34, negated_conjecture, (k6_mcart_1(esk4_0,esk5_0,esk6_0,esk7_0)!=k6_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_32])])).
cnf(c_0_35, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_33]), c_0_34]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 36
# Proof object clause steps            : 29
# Proof object formula steps           : 7
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 35
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 14
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 13
# Processed clauses                    : 35
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 34
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 6
# ...of the previous two non-trivial   : 9
# Contextual simplify-reflections      : 0
# Paramodulations                      : 6
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 18
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 3
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 17
# Clause-clause subsumption calls (NU) : 6
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 8
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 923

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.000 s
# Total time               : 0.009 s
# Maximum resident set size: 28680 pages
