# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.019 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d21_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X3=k15_lattice3(X1,X2)<=>(r4_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r4_lattice3(X1,X4,X2)=>r1_lattices(X1,X3,X4)))))))), file('data/MPTP/proofs/3d51f818cb0542048f011e7a763fcf4f.E_input', d21_lattice3)).
fof(dt_k15_lattice3, axiom, ![X1, X2]:((~(v2_struct_0(X1))&l3_lattices(X1))=>m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/3d51f818cb0542048f011e7a763fcf4f.E_input', dt_k15_lattice3)).
fof(t45_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(X2=k15_lattice3(X1,a_2_3_lattice3(X1,X2))&X2=k16_lattice3(X1,a_2_4_lattice3(X1,X2))))), file('data/MPTP/proofs/3d51f818cb0542048f011e7a763fcf4f.E_input', t45_lattice3)).
fof(d17_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r4_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X4,X2)))))), file('data/MPTP/proofs/3d51f818cb0542048f011e7a763fcf4f.E_input', d17_lattice3)).
fof(fraenkel_a_2_3_lattice3, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))&m1_subset_1(X3,u1_struct_0(X2)))=>(r2_hidden(X1,a_2_3_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&r3_lattices(X2,X4,X3))))), file('data/MPTP/proofs/3d51f818cb0542048f011e7a763fcf4f.E_input', fraenkel_a_2_3_lattice3)).
fof(t50_lattice3, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))&k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0))), file('data/MPTP/proofs/3d51f818cb0542048f011e7a763fcf4f.E_input', t50_lattice3)).
fof(t48_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X2)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>~((r3_lattices(X1,X4,X5)&r2_hidden(X5,X3)))))))=>r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3)))), file('data/MPTP/proofs/3d51f818cb0542048f011e7a763fcf4f.E_input', t48_lattice3)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/3d51f818cb0542048f011e7a763fcf4f.E_input', t113_zfmisc_1)).
fof(t152_zfmisc_1, axiom, ![X1, X2]:~(r2_hidden(X1,k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/3d51f818cb0542048f011e7a763fcf4f.E_input', t152_zfmisc_1)).
fof(d9_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>(v9_lattices(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_lattices(X1,X2,k1_lattices(X1,X2,X3))=X2)))), file('data/MPTP/proofs/3d51f818cb0542048f011e7a763fcf4f.E_input', d9_lattices)).
fof(d3_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l2_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattices(X1,X2,X3)<=>k1_lattices(X1,X2,X3)=X3)))), file('data/MPTP/proofs/3d51f818cb0542048f011e7a763fcf4f.E_input', d3_lattices)).
fof(dt_l3_lattices, axiom, ![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1))), file('data/MPTP/proofs/3d51f818cb0542048f011e7a763fcf4f.E_input', dt_l3_lattices)).
fof(d16_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v13_lattices(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(X2=k5_lattices(X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_lattices(X1,X2,X3)=X2&k2_lattices(X1,X3,X2)=X2)))))), file('data/MPTP/proofs/3d51f818cb0542048f011e7a763fcf4f.E_input', d16_lattices)).
fof(dt_k5_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>m1_subset_1(k5_lattices(X1),u1_struct_0(X1))), file('data/MPTP/proofs/3d51f818cb0542048f011e7a763fcf4f.E_input', dt_k5_lattices)).
fof(d13_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v13_lattices(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_lattices(X1,X2,X3)=X2&k2_lattices(X1,X3,X2)=X2))))), file('data/MPTP/proofs/3d51f818cb0542048f011e7a763fcf4f.E_input', d13_lattices)).
fof(d6_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v6_lattices(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2))))), file('data/MPTP/proofs/3d51f818cb0542048f011e7a763fcf4f.E_input', d6_lattices)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/3d51f818cb0542048f011e7a763fcf4f.E_input', cc1_lattices)).
fof(c_0_17, plain, ![X6, X7, X8, X9]:(((r4_lattice3(X6,X8,X7)|X8!=k15_lattice3(X6,X7)|~m1_subset_1(X8,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6))))|(v2_struct_0(X6)|~l3_lattices(X6)))&(~m1_subset_1(X9,u1_struct_0(X6))|(~r4_lattice3(X6,X9,X7)|r1_lattices(X6,X8,X9))|X8!=k15_lattice3(X6,X7)|~m1_subset_1(X8,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6))))|(v2_struct_0(X6)|~l3_lattices(X6))))&((m1_subset_1(esk1_3(X6,X7,X8),u1_struct_0(X6))|~r4_lattice3(X6,X8,X7)|X8=k15_lattice3(X6,X7)|~m1_subset_1(X8,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6))))|(v2_struct_0(X6)|~l3_lattices(X6)))&((r4_lattice3(X6,esk1_3(X6,X7,X8),X7)|~r4_lattice3(X6,X8,X7)|X8=k15_lattice3(X6,X7)|~m1_subset_1(X8,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6))))|(v2_struct_0(X6)|~l3_lattices(X6)))&(~r1_lattices(X6,X8,esk1_3(X6,X7,X8))|~r4_lattice3(X6,X8,X7)|X8=k15_lattice3(X6,X7)|~m1_subset_1(X8,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6))))|(v2_struct_0(X6)|~l3_lattices(X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d21_lattice3])])])])])])).
cnf(c_0_18, plain, (r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X1)|X2!=k15_lattice3(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_19, plain, ![X11, X12]:(v2_struct_0(X11)|~l3_lattices(X11)|m1_subset_1(k15_lattice3(X11,X12),u1_struct_0(X11))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k15_lattice3])])])).
cnf(c_0_20, plain, (v2_struct_0(X1)|r4_lattice3(X1,X2,X3)|X2!=k15_lattice3(X1,X3)|~l3_lattices(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_18])).
cnf(c_0_21, plain, (v2_struct_0(X1)|m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_22, plain, ![X13, X14]:((X14=k15_lattice3(X13,a_2_3_lattice3(X13,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(v2_struct_0(X13)|(~v10_lattices(X13)|(~v4_lattice3(X13)|~l3_lattices(X13)))))&(X14=k16_lattice3(X13,a_2_4_lattice3(X13,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(v2_struct_0(X13)|(~v10_lattices(X13)|(~v4_lattice3(X13)|~l3_lattices(X13)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t45_lattice3])])])])])).
fof(c_0_23, plain, ![X15, X16, X17, X18, X19]:((~r4_lattice3(X15,X16,X17)|(~m1_subset_1(X18,u1_struct_0(X15))|(~r2_hidden(X18,X17)|r1_lattices(X15,X18,X16)))|~m1_subset_1(X16,u1_struct_0(X15))|(v2_struct_0(X15)|~l3_lattices(X15)))&((m1_subset_1(esk2_3(X15,X16,X19),u1_struct_0(X15))|r4_lattice3(X15,X16,X19)|~m1_subset_1(X16,u1_struct_0(X15))|(v2_struct_0(X15)|~l3_lattices(X15)))&((r2_hidden(esk2_3(X15,X16,X19),X19)|r4_lattice3(X15,X16,X19)|~m1_subset_1(X16,u1_struct_0(X15))|(v2_struct_0(X15)|~l3_lattices(X15)))&(~r1_lattices(X15,esk2_3(X15,X16,X19),X16)|r4_lattice3(X15,X16,X19)|~m1_subset_1(X16,u1_struct_0(X15))|(v2_struct_0(X15)|~l3_lattices(X15)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_lattice3])])])])])])])).
cnf(c_0_24, plain, (r4_lattice3(X1,k15_lattice3(X1,X2),X2)|v2_struct_0(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_20]), c_0_21])).
cnf(c_0_25, plain, (X1=k15_lattice3(X2,a_2_3_lattice3(X2,X1))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_26, plain, ![X21, X22, X23, X25]:(((m1_subset_1(esk3_3(X21,X22,X23),u1_struct_0(X22))|~r2_hidden(X21,a_2_3_lattice3(X22,X23))|(v2_struct_0(X22)|(~v10_lattices(X22)|(~v4_lattice3(X22)|~l3_lattices(X22)))|~m1_subset_1(X23,u1_struct_0(X22))))&((X21=esk3_3(X21,X22,X23)|~r2_hidden(X21,a_2_3_lattice3(X22,X23))|(v2_struct_0(X22)|(~v10_lattices(X22)|(~v4_lattice3(X22)|~l3_lattices(X22)))|~m1_subset_1(X23,u1_struct_0(X22))))&(r3_lattices(X22,esk3_3(X21,X22,X23),X23)|~r2_hidden(X21,a_2_3_lattice3(X22,X23))|(v2_struct_0(X22)|(~v10_lattices(X22)|(~v4_lattice3(X22)|~l3_lattices(X22)))|~m1_subset_1(X23,u1_struct_0(X22))))))&(~m1_subset_1(X25,u1_struct_0(X22))|(X21!=X25|~r3_lattices(X22,X25,X23))|r2_hidden(X21,a_2_3_lattice3(X22,X23))|(v2_struct_0(X22)|(~v10_lattices(X22)|(~v4_lattice3(X22)|~l3_lattices(X22)))|~m1_subset_1(X23,u1_struct_0(X22))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_3_lattice3])])])])])])).
fof(c_0_27, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))&k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0)))), inference(assume_negation,[status(cth)],[t50_lattice3])).
cnf(c_0_28, plain, (r1_lattices(X1,X4,X2)|v2_struct_0(X1)|~r4_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (r4_lattice3(X1,X2,a_2_3_lattice3(X1,X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_30, plain, (r2_hidden(X3,a_2_3_lattice3(X2,X4))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~r3_lattices(X2,X1,X4)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_31, plain, ![X27, X28, X29, X31]:((m1_subset_1(esk5_3(X27,X28,X29),u1_struct_0(X27))|r3_lattices(X27,k15_lattice3(X27,X28),k15_lattice3(X27,X29))|(v2_struct_0(X27)|(~v10_lattices(X27)|(~v4_lattice3(X27)|~l3_lattices(X27)))))&((r2_hidden(esk5_3(X27,X28,X29),X28)|r3_lattices(X27,k15_lattice3(X27,X28),k15_lattice3(X27,X29))|(v2_struct_0(X27)|(~v10_lattices(X27)|(~v4_lattice3(X27)|~l3_lattices(X27)))))&(~m1_subset_1(X31,u1_struct_0(X27))|(~r3_lattices(X27,esk5_3(X27,X28,X29),X31)|~r2_hidden(X31,X29))|r3_lattices(X27,k15_lattice3(X27,X28),k15_lattice3(X27,X29))|(v2_struct_0(X27)|(~v10_lattices(X27)|(~v4_lattice3(X27)|~l3_lattices(X27))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t48_lattice3])])])])])])).
fof(c_0_32, negated_conjecture, ((~v2_struct_0(esk4_0)&(v10_lattices(esk4_0)&(v4_lattice3(esk4_0)&l3_lattices(esk4_0))))&(v2_struct_0(esk4_0)|(~v10_lattices(esk4_0)|(~v13_lattices(esk4_0)|~l3_lattices(esk4_0)))|k5_lattices(esk4_0)!=k15_lattice3(esk4_0,k1_xboole_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_27])])])])).
cnf(c_0_33, plain, (r1_lattices(X1,X2,X3)|v2_struct_0(X1)|~r2_hidden(X2,a_2_3_lattice3(X1,X3))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34, plain, (r2_hidden(X1,a_2_3_lattice3(X2,X3))|v2_struct_0(X2)|~r3_lattices(X2,X1,X3)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~v4_lattice3(X2)|~v10_lattices(X2)|~l3_lattices(X2)), inference(er,[status(thm)],[c_0_30])).
cnf(c_0_35, plain, (r2_hidden(esk5_3(X1,X2,X3),X2)|r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3))|v2_struct_0(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_36, negated_conjecture, (v4_lattice3(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_37, negated_conjecture, (v10_lattices(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, negated_conjecture, (l3_lattices(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40, plain, (r1_lattices(X1,X2,X3)|v2_struct_0(X1)|~r3_lattices(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_41, negated_conjecture, (r3_lattices(esk4_0,k15_lattice3(esk4_0,X1),k15_lattice3(esk4_0,X2))|r2_hidden(esk5_3(esk4_0,X1,X2),X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37]), c_0_38])]), c_0_39])).
fof(c_0_42, plain, ![X32, X33]:((k2_zfmisc_1(X32,X33)!=k1_xboole_0|(X32=k1_xboole_0|X33=k1_xboole_0))&((X32!=k1_xboole_0|k2_zfmisc_1(X32,X33)=k1_xboole_0)&(X33!=k1_xboole_0|k2_zfmisc_1(X32,X33)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_43, negated_conjecture, (r2_hidden(esk5_3(esk4_0,X1,X2),X1)|r1_lattices(esk4_0,k15_lattice3(esk4_0,X1),k15_lattice3(esk4_0,X2))|~m1_subset_1(k15_lattice3(esk4_0,X1),u1_struct_0(esk4_0))|~m1_subset_1(k15_lattice3(esk4_0,X2),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_36]), c_0_37]), c_0_38])]), c_0_39])).
fof(c_0_44, plain, ![X34, X35]:~r2_hidden(X34,k2_zfmisc_1(X34,X35)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_zfmisc_1])])).
cnf(c_0_45, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk5_3(esk4_0,X1,X2),X1)|r1_lattices(esk4_0,k15_lattice3(esk4_0,X1),k15_lattice3(esk4_0,X2))|~m1_subset_1(k15_lattice3(esk4_0,X2),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_21]), c_0_38])]), c_0_39])).
fof(c_0_47, plain, ![X36, X37, X38]:((~v9_lattices(X36)|(~m1_subset_1(X37,u1_struct_0(X36))|(~m1_subset_1(X38,u1_struct_0(X36))|k2_lattices(X36,X37,k1_lattices(X36,X37,X38))=X37))|(v2_struct_0(X36)|~l3_lattices(X36)))&((m1_subset_1(esk6_1(X36),u1_struct_0(X36))|v9_lattices(X36)|(v2_struct_0(X36)|~l3_lattices(X36)))&((m1_subset_1(esk7_1(X36),u1_struct_0(X36))|v9_lattices(X36)|(v2_struct_0(X36)|~l3_lattices(X36)))&(k2_lattices(X36,esk6_1(X36),k1_lattices(X36,esk6_1(X36),esk7_1(X36)))!=esk6_1(X36)|v9_lattices(X36)|(v2_struct_0(X36)|~l3_lattices(X36)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d9_lattices])])])])])])).
fof(c_0_48, plain, ![X41, X42, X43]:((~r1_lattices(X41,X42,X43)|k1_lattices(X41,X42,X43)=X43|~m1_subset_1(X43,u1_struct_0(X41))|~m1_subset_1(X42,u1_struct_0(X41))|(v2_struct_0(X41)|~l2_lattices(X41)))&(k1_lattices(X41,X42,X43)!=X43|r1_lattices(X41,X42,X43)|~m1_subset_1(X43,u1_struct_0(X41))|~m1_subset_1(X42,u1_struct_0(X41))|(v2_struct_0(X41)|~l2_lattices(X41)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d3_lattices])])])])])).
fof(c_0_49, plain, ![X44]:((l1_lattices(X44)|~l3_lattices(X44))&(l2_lattices(X44)|~l3_lattices(X44))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
cnf(c_0_50, plain, (~r2_hidden(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_51, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_45])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk5_3(esk4_0,X1,X2),X1)|r1_lattices(esk4_0,k15_lattice3(esk4_0,X1),k15_lattice3(esk4_0,X2))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_21]), c_0_38])]), c_0_39])).
cnf(c_0_53, plain, (k2_lattices(X1,X2,k1_lattices(X1,X2,X3))=X2|v2_struct_0(X1)|~v9_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_54, plain, (k1_lattices(X1,X2,X3)=X3|v2_struct_0(X1)|~r1_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_55, plain, (l2_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_56, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_57, negated_conjecture, (r2_hidden(esk5_3(esk4_0,X1,a_2_3_lattice3(esk4_0,X2)),X1)|r1_lattices(esk4_0,k15_lattice3(esk4_0,X1),X2)|~m1_subset_1(X2,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_25]), c_0_36]), c_0_37]), c_0_38])]), c_0_39])).
fof(c_0_58, plain, ![X45, X46, X47]:(((k2_lattices(X45,X46,X47)=X46|~m1_subset_1(X47,u1_struct_0(X45))|X46!=k5_lattices(X45)|~m1_subset_1(X46,u1_struct_0(X45))|~v13_lattices(X45)|(v2_struct_0(X45)|~l1_lattices(X45)))&(k2_lattices(X45,X47,X46)=X46|~m1_subset_1(X47,u1_struct_0(X45))|X46!=k5_lattices(X45)|~m1_subset_1(X46,u1_struct_0(X45))|~v13_lattices(X45)|(v2_struct_0(X45)|~l1_lattices(X45))))&((m1_subset_1(esk8_2(X45,X46),u1_struct_0(X45))|X46=k5_lattices(X45)|~m1_subset_1(X46,u1_struct_0(X45))|~v13_lattices(X45)|(v2_struct_0(X45)|~l1_lattices(X45)))&(k2_lattices(X45,X46,esk8_2(X45,X46))!=X46|k2_lattices(X45,esk8_2(X45,X46),X46)!=X46|X46=k5_lattices(X45)|~m1_subset_1(X46,u1_struct_0(X45))|~v13_lattices(X45)|(v2_struct_0(X45)|~l1_lattices(X45))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattices])])])])])])).
fof(c_0_59, plain, ![X49]:(v2_struct_0(X49)|~l1_lattices(X49)|m1_subset_1(k5_lattices(X49),u1_struct_0(X49))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_lattices])])])).
fof(c_0_60, plain, ![X56, X58, X59]:(((m1_subset_1(esk11_1(X56),u1_struct_0(X56))|~v13_lattices(X56)|(v2_struct_0(X56)|~l1_lattices(X56)))&((k2_lattices(X56,esk11_1(X56),X58)=esk11_1(X56)|~m1_subset_1(X58,u1_struct_0(X56))|~v13_lattices(X56)|(v2_struct_0(X56)|~l1_lattices(X56)))&(k2_lattices(X56,X58,esk11_1(X56))=esk11_1(X56)|~m1_subset_1(X58,u1_struct_0(X56))|~v13_lattices(X56)|(v2_struct_0(X56)|~l1_lattices(X56)))))&((m1_subset_1(esk12_2(X56,X59),u1_struct_0(X56))|~m1_subset_1(X59,u1_struct_0(X56))|v13_lattices(X56)|(v2_struct_0(X56)|~l1_lattices(X56)))&(k2_lattices(X56,X59,esk12_2(X56,X59))!=X59|k2_lattices(X56,esk12_2(X56,X59),X59)!=X59|~m1_subset_1(X59,u1_struct_0(X56))|v13_lattices(X56)|(v2_struct_0(X56)|~l1_lattices(X56))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_lattices])])])])])])).
cnf(c_0_61, plain, (k2_lattices(X1,X2,X3)=X2|v2_struct_0(X1)|~v9_lattices(X1)|~r1_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55])).
cnf(c_0_62, negated_conjecture, (r1_lattices(esk4_0,k15_lattice3(esk4_0,k1_xboole_0),X1)|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_63, plain, (k2_lattices(X1,X2,X3)=X2|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|X2!=k5_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v13_lattices(X1)|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_64, plain, (v2_struct_0(X1)|m1_subset_1(k5_lattices(X1),u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_65, plain, (k2_lattices(X1,X2,esk11_1(X1))=esk11_1(X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v13_lattices(X1)|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_66, negated_conjecture, (k2_lattices(esk4_0,k15_lattice3(esk4_0,k1_xboole_0),X1)=k15_lattice3(esk4_0,k1_xboole_0)|~v9_lattices(esk4_0)|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_38])]), c_0_39])).
cnf(c_0_67, plain, (k2_lattices(X1,k5_lattices(X1),X2)=k5_lattices(X1)|v2_struct_0(X1)|~l1_lattices(X1)|~v13_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_63]), c_0_64])).
cnf(c_0_68, plain, (m1_subset_1(esk11_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|~v13_lattices(X1)|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_69, negated_conjecture, (esk11_1(esk4_0)=k15_lattice3(esk4_0,k1_xboole_0)|~l1_lattices(esk4_0)|~v9_lattices(esk4_0)|~v13_lattices(esk4_0)|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))|~m1_subset_1(esk11_1(esk4_0),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_39])).
cnf(c_0_70, negated_conjecture, (v2_struct_0(esk4_0)|~v10_lattices(esk4_0)|~v13_lattices(esk4_0)|~l3_lattices(esk4_0)|k5_lattices(esk4_0)!=k15_lattice3(esk4_0,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_71, plain, ![X50, X51, X52]:((~v6_lattices(X50)|(~m1_subset_1(X51,u1_struct_0(X50))|(~m1_subset_1(X52,u1_struct_0(X50))|k2_lattices(X50,X51,X52)=k2_lattices(X50,X52,X51)))|(v2_struct_0(X50)|~l1_lattices(X50)))&((m1_subset_1(esk9_1(X50),u1_struct_0(X50))|v6_lattices(X50)|(v2_struct_0(X50)|~l1_lattices(X50)))&((m1_subset_1(esk10_1(X50),u1_struct_0(X50))|v6_lattices(X50)|(v2_struct_0(X50)|~l1_lattices(X50)))&(k2_lattices(X50,esk9_1(X50),esk10_1(X50))!=k2_lattices(X50,esk10_1(X50),esk9_1(X50))|v6_lattices(X50)|(v2_struct_0(X50)|~l1_lattices(X50)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_lattices])])])])])])).
fof(c_0_72, plain, ![X55]:((~v2_struct_0(X55)|(v2_struct_0(X55)|~v10_lattices(X55))|~l3_lattices(X55))&((v4_lattices(X55)|(v2_struct_0(X55)|~v10_lattices(X55))|~l3_lattices(X55))&((v5_lattices(X55)|(v2_struct_0(X55)|~v10_lattices(X55))|~l3_lattices(X55))&((v6_lattices(X55)|(v2_struct_0(X55)|~v10_lattices(X55))|~l3_lattices(X55))&((v7_lattices(X55)|(v2_struct_0(X55)|~v10_lattices(X55))|~l3_lattices(X55))&((v8_lattices(X55)|(v2_struct_0(X55)|~v10_lattices(X55))|~l3_lattices(X55))&(v9_lattices(X55)|(v2_struct_0(X55)|~v10_lattices(X55))|~l3_lattices(X55)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
cnf(c_0_73, plain, (esk11_1(X1)=k5_lattices(X1)|v2_struct_0(X1)|~l1_lattices(X1)|~v13_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_65]), c_0_64]), c_0_68])).
cnf(c_0_74, negated_conjecture, (esk11_1(esk4_0)=k15_lattice3(esk4_0,k1_xboole_0)|~l1_lattices(esk4_0)|~v9_lattices(esk4_0)|~v13_lattices(esk4_0)|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_68]), c_0_39])).
cnf(c_0_75, negated_conjecture, (k5_lattices(esk4_0)!=k15_lattice3(esk4_0,k1_xboole_0)|~v13_lattices(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_38]), c_0_37])]), c_0_39])).
cnf(c_0_76, plain, (k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2)|v2_struct_0(X1)|~v6_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_77, plain, (v6_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_78, plain, (l1_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_79, negated_conjecture, (~l1_lattices(esk4_0)|~v9_lattices(esk4_0)|~v13_lattices(esk4_0)|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_39]), c_0_75])).
cnf(c_0_80, plain, (v9_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_81, plain, (k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78])).
cnf(c_0_82, negated_conjecture, (~l1_lattices(esk4_0)|~v13_lattices(esk4_0)|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_83, plain, (v13_lattices(X1)|v2_struct_0(X1)|k2_lattices(X1,X2,esk12_2(X1,X2))!=X2|k2_lattices(X1,esk12_2(X1,X2),X2)!=X2|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_84, negated_conjecture, (k2_lattices(esk4_0,X1,k15_lattice3(esk4_0,k1_xboole_0))=k15_lattice3(esk4_0,k1_xboole_0)|~v9_lattices(esk4_0)|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_81]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_85, negated_conjecture, (~l1_lattices(esk4_0)|~v13_lattices(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_21]), c_0_38])]), c_0_39])).
cnf(c_0_86, negated_conjecture, (~l1_lattices(esk4_0)|~v9_lattices(esk4_0)|~m1_subset_1(esk12_2(esk4_0,k15_lattice3(esk4_0,k1_xboole_0)),u1_struct_0(esk4_0))|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_39]), c_0_66]), c_0_85])).
cnf(c_0_87, plain, (m1_subset_1(esk12_2(X1,X2),u1_struct_0(X1))|v13_lattices(X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_88, negated_conjecture, (~l1_lattices(esk4_0)|~v9_lattices(esk4_0)|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_39]), c_0_85])).
cnf(c_0_89, negated_conjecture, (~l1_lattices(esk4_0)|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_80]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_90, negated_conjecture, (~l1_lattices(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_21]), c_0_38])]), c_0_39])).
cnf(c_0_91, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_78]), c_0_38])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 92
# Proof object clause steps            : 57
# Proof object formula steps           : 35
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 67
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 57
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 56
# Processed clauses                    : 218
# ...of these trivial                  : 0
# ...subsumed                          : 32
# ...remaining for further processing  : 186
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 12
# Backward-rewritten                   : 0
# Generated clauses                    : 195
# ...of the previous two non-trivial   : 184
# Contextual simplify-reflections      : 19
# Paramodulations                      : 188
# Factorizations                       : 0
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 111
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 102
# Current number of unprocessed clauses: 72
# ...number of literals in the above   : 544
# Current number of archived formulas  : 0
# Current number of archived clauses   : 68
# Clause-clause subsumption calls (NU) : 4533
# Rec. Clause-clause subsumption calls : 391
# Non-unit clause-clause subsumptions  : 61
# Unit Clause-clause subsumption calls : 54
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 14310

# -------------------------------------------------
# User time                : 0.030 s
# System time              : 0.000 s
# Total time               : 0.030 s
# Maximum resident set size: 28836 pages
