# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t5_xboole_1, axiom, ![X1, X2, X3]:k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_xboole_0(X1,X3),k2_xboole_0(X2,X3)), file('data/MPTP/proofs/3d7075988fae4ce2a3abb15353dc01bc.E_input', t5_xboole_1)).
fof(t4_xboole_1, axiom, ![X1, X2, X3]:k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/3d7075988fae4ce2a3abb15353dc01bc.E_input', t4_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/3d7075988fae4ce2a3abb15353dc01bc.E_input', commutativity_k2_xboole_0)).
fof(t50_enumset1, axiom, ![X1, X2, X3, X4, X5]:k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_tarski(X5)), file('data/MPTP/proofs/3d7075988fae4ce2a3abb15353dc01bc.E_input', t50_enumset1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/3d7075988fae4ce2a3abb15353dc01bc.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/3d7075988fae4ce2a3abb15353dc01bc.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/3d7075988fae4ce2a3abb15353dc01bc.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/3d7075988fae4ce2a3abb15353dc01bc.E_input', t72_enumset1)).
fof(t47_enumset1, axiom, ![X1, X2, X3, X4, X5]:k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k1_tarski(X1),k2_enumset1(X2,X3,X4,X5)), file('data/MPTP/proofs/3d7075988fae4ce2a3abb15353dc01bc.E_input', t47_enumset1)).
fof(t48_enumset1, axiom, ![X1, X2, X3, X4, X5]:k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_tarski(X1,X2),k1_enumset1(X3,X4,X5)), file('data/MPTP/proofs/3d7075988fae4ce2a3abb15353dc01bc.E_input', t48_enumset1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/3d7075988fae4ce2a3abb15353dc01bc.E_input', commutativity_k2_tarski)).
fof(t73_enumset1, conjecture, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/3d7075988fae4ce2a3abb15353dc01bc.E_input', t73_enumset1)).
fof(t51_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k1_tarski(X1),k3_enumset1(X2,X3,X4,X5,X6)), file('data/MPTP/proofs/3d7075988fae4ce2a3abb15353dc01bc.E_input', t51_enumset1)).
fof(c_0_13, plain, ![X7, X8, X9]:k2_xboole_0(k2_xboole_0(X7,X8),X9)=k2_xboole_0(k2_xboole_0(X7,X9),k2_xboole_0(X8,X9)), inference(variable_rename,[status(thm)],[t5_xboole_1])).
fof(c_0_14, plain, ![X10, X11, X12]:k2_xboole_0(k2_xboole_0(X10,X11),X12)=k2_xboole_0(X10,k2_xboole_0(X11,X12)), inference(variable_rename,[status(thm)],[t4_xboole_1])).
fof(c_0_15, plain, ![X13, X14]:k2_xboole_0(X13,X14)=k2_xboole_0(X14,X13), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_16, plain, (k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_xboole_0(X1,X3),k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, plain, (k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, plain, (k2_xboole_0(X1,k2_xboole_0(X2,k2_xboole_0(X3,X2)))=k2_xboole_0(X1,k2_xboole_0(X3,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_17]), c_0_17])).
fof(c_0_20, plain, ![X15, X16, X17, X18, X19]:k3_enumset1(X15,X16,X17,X18,X19)=k2_xboole_0(k2_enumset1(X15,X16,X17,X18),k1_tarski(X19)), inference(variable_rename,[status(thm)],[t50_enumset1])).
fof(c_0_21, plain, ![X20]:k2_tarski(X20,X20)=k1_tarski(X20), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_22, plain, ![X21, X22]:k1_enumset1(X21,X21,X22)=k2_tarski(X21,X22), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_23, plain, ![X23, X24, X25]:k2_enumset1(X23,X23,X24,X25)=k1_enumset1(X23,X24,X25), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_24, plain, ![X26, X27, X28, X29]:k3_enumset1(X26,X26,X27,X28,X29)=k2_enumset1(X26,X27,X28,X29), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_25, plain, ![X30, X31, X32, X33, X34]:k3_enumset1(X30,X31,X32,X33,X34)=k2_xboole_0(k1_tarski(X30),k2_enumset1(X31,X32,X33,X34)), inference(variable_rename,[status(thm)],[t47_enumset1])).
cnf(c_0_26, plain, (k2_xboole_0(X1,k2_xboole_0(X2,X3))=k2_xboole_0(X3,k2_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_18, c_0_17])).
cnf(c_0_27, plain, (k2_xboole_0(X1,k2_xboole_0(X2,k2_xboole_0(X1,X3)))=k2_xboole_0(X3,k2_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_17]), c_0_17])).
fof(c_0_28, plain, ![X35, X36, X37, X38, X39]:k3_enumset1(X35,X36,X37,X38,X39)=k2_xboole_0(k2_tarski(X35,X36),k1_enumset1(X37,X38,X39)), inference(variable_rename,[status(thm)],[t48_enumset1])).
fof(c_0_29, plain, ![X40, X41]:k2_tarski(X40,X41)=k2_tarski(X41,X40), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_30, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_tarski(X5))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_31, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_32, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_34, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_35, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k1_tarski(X1),k2_enumset1(X2,X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_36, negated_conjecture, ~(![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(assume_negation,[status(cth)],[t73_enumset1])).
fof(c_0_37, plain, ![X47, X48, X49, X50, X51, X52]:k4_enumset1(X47,X48,X49,X50,X51,X52)=k2_xboole_0(k1_tarski(X47),k3_enumset1(X48,X49,X50,X51,X52)), inference(variable_rename,[status(thm)],[t51_enumset1])).
cnf(c_0_38, plain, (k2_xboole_0(X1,k2_xboole_0(X2,k2_xboole_0(X3,X3)))=k2_xboole_0(X1,k2_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_19, c_0_26])).
cnf(c_0_39, plain, (k2_xboole_0(X1,k2_xboole_0(X1,X2))=k2_xboole_0(X2,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_19, c_0_27])).
cnf(c_0_40, plain, (k2_xboole_0(X1,k2_xboole_0(X1,k2_xboole_0(X2,X3)))=k2_xboole_0(X2,k2_xboole_0(X3,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_18]), c_0_17])).
cnf(c_0_41, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_tarski(X1,X2),k1_enumset1(X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_42, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_43, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k3_enumset1(X1,X1,X2,X3,X4),k3_enumset1(X5,X5,X5,X5,X5))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_34])).
cnf(c_0_44, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k3_enumset1(X1,X1,X1,X1,X1),k3_enumset1(X2,X2,X3,X4,X5))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_34])).
fof(c_0_45, negated_conjecture, k4_enumset1(esk1_0,esk1_0,esk2_0,esk3_0,esk4_0,esk5_0)!=k3_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_36])])])).
cnf(c_0_46, plain, (k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k1_tarski(X1),k3_enumset1(X2,X3,X4,X5,X6))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_47, plain, (k2_xboole_0(X1,k2_xboole_0(X2,X2))=k2_xboole_0(X1,k2_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_17]), c_0_17]), c_0_40]), c_0_40])).
cnf(c_0_48, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k3_enumset1(X1,X1,X1,X1,X2),k3_enumset1(X3,X3,X3,X4,X5))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_32]), c_0_33]), c_0_33]), c_0_34]), c_0_34])).
cnf(c_0_49, plain, (k3_enumset1(X1,X1,X1,X1,X2)=k3_enumset1(X2,X2,X2,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_32]), c_0_32]), c_0_33]), c_0_33]), c_0_34]), c_0_34])).
cnf(c_0_50, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k3_enumset1(X2,X3,X4,X5,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_43]), c_0_44])).
cnf(c_0_51, negated_conjecture, (k4_enumset1(esk1_0,esk1_0,esk2_0,esk3_0,esk4_0,esk5_0)!=k3_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_52, plain, (k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k3_enumset1(X1,X1,X1,X1,X1),k3_enumset1(X2,X3,X4,X5,X6))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_31]), c_0_32]), c_0_33]), c_0_34])).
cnf(c_0_53, plain, (k2_xboole_0(k3_enumset1(X1,X1,X2,X3,X4),k3_enumset1(X1,X2,X3,X4,X5))=k2_xboole_0(k3_enumset1(X5,X5,X5,X5,X5),k3_enumset1(X5,X1,X2,X3,X4))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_44]), c_0_43])).
cnf(c_0_54, plain, (k2_xboole_0(k3_enumset1(X1,X1,X2,X3,X4),k3_enumset1(X1,X2,X3,X4,X5))=k3_enumset1(X1,X2,X3,X4,X5)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_43]), c_0_44]), c_0_43])).
cnf(c_0_55, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k3_enumset1(X2,X1,X3,X4,X5)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_48])).
cnf(c_0_56, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k3_enumset1(X3,X4,X5,X1,X2)), inference(spm,[status(thm)],[c_0_50, c_0_50])).
cnf(c_0_57, negated_conjecture, (k2_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0))!=k3_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0)), inference(rw,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_58, plain, (k2_xboole_0(k3_enumset1(X1,X1,X1,X1,X1),k3_enumset1(X1,X2,X3,X4,X5))=k3_enumset1(X2,X3,X4,X5,X1)), inference(rw,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_59, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k3_enumset1(X5,X4,X1,X2,X3)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_60, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k3_enumset1(X1,X3,X4,X5,X2)), inference(spm,[status(thm)],[c_0_50, c_0_55])).
cnf(c_0_61, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_58]), c_0_59]), c_0_60])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 62
# Proof object clause steps            : 35
# Proof object formula steps           : 27
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 45
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 13
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 8
# Processed clauses                    : 1403
# ...of these trivial                  : 334
# ...subsumed                          : 941
# ...remaining for further processing  : 128
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 6
# Generated clauses                    : 20355
# ...of the previous two non-trivial   : 17305
# Contextual simplify-reflections      : 0
# Paramodulations                      : 20355
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 110
#    Positive orientable unit clauses  : 36
#    Positive unorientable unit clauses: 74
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 15916
# ...number of literals in the above   : 15916
# Current number of archived formulas  : 0
# Current number of archived clauses   : 23
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 3425
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8111
# BW rewrite match successes           : 2507
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 116966

# -------------------------------------------------
# User time                : 0.094 s
# System time              : 0.000 s
# Total time               : 0.094 s
# Maximum resident set size: 28760 pages
