# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t48_funct_1, conjecture, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((v2_funct_1(k5_relat_1(X2,X1))&k2_relat_1(X2)=k1_relat_1(X1))=>(v2_funct_1(X2)&v2_funct_1(X1))))), file('data/MPTP/proofs/3d72a676786944429139202523c149c7.E_input', t48_funct_1)).
fof(t47_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((v2_funct_1(k5_relat_1(X2,X1))&r1_tarski(k2_relat_1(X2),k1_relat_1(X1)))=>v2_funct_1(X2)))), file('data/MPTP/proofs/3d72a676786944429139202523c149c7.E_input', t47_funct_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/3d72a676786944429139202523c149c7.E_input', d10_xboole_0)).
fof(d8_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)<=>![X2, X3]:((r2_hidden(X2,k1_relat_1(X1))&(r2_hidden(X3,k1_relat_1(X1))&k1_funct_1(X1,X2)=k1_funct_1(X1,X3)))=>X2=X3))), file('data/MPTP/proofs/3d72a676786944429139202523c149c7.E_input', d8_funct_1)).
fof(d5_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4))))), file('data/MPTP/proofs/3d72a676786944429139202523c149c7.E_input', d5_funct_1)).
fof(t46_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(k2_relat_1(X1),k1_relat_1(X2))=>k1_relat_1(k5_relat_1(X1,X2))=k1_relat_1(X1)))), file('data/MPTP/proofs/3d72a676786944429139202523c149c7.E_input', t46_relat_1)).
fof(fc2_funct_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&(v1_relat_1(X2)&v1_funct_1(X2)))=>(v1_relat_1(k5_relat_1(X1,X2))&v1_funct_1(k5_relat_1(X1,X2)))), file('data/MPTP/proofs/3d72a676786944429139202523c149c7.E_input', fc2_funct_1)).
fof(dt_k5_relat_1, axiom, ![X1, X2]:((v1_relat_1(X1)&v1_relat_1(X2))=>v1_relat_1(k5_relat_1(X1,X2))), file('data/MPTP/proofs/3d72a676786944429139202523c149c7.E_input', dt_k5_relat_1)).
fof(t23_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k1_relat_1(X2))=>k1_funct_1(k5_relat_1(X2,X3),X1)=k1_funct_1(X3,k1_funct_1(X2,X1))))), file('data/MPTP/proofs/3d72a676786944429139202523c149c7.E_input', t23_funct_1)).
fof(c_0_9, negated_conjecture, ~(![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((v2_funct_1(k5_relat_1(X2,X1))&k2_relat_1(X2)=k1_relat_1(X1))=>(v2_funct_1(X2)&v2_funct_1(X1)))))), inference(assume_negation,[status(cth)],[t48_funct_1])).
fof(c_0_10, plain, ![X7, X8]:(~v1_relat_1(X7)|~v1_funct_1(X7)|(~v1_relat_1(X8)|~v1_funct_1(X8)|(~v2_funct_1(k5_relat_1(X8,X7))|~r1_tarski(k2_relat_1(X8),k1_relat_1(X7))|v2_funct_1(X8)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t47_funct_1])])])).
fof(c_0_11, negated_conjecture, ((v1_relat_1(esk1_0)&v1_funct_1(esk1_0))&((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&((v2_funct_1(k5_relat_1(esk2_0,esk1_0))&k2_relat_1(esk2_0)=k1_relat_1(esk1_0))&(~v2_funct_1(esk2_0)|~v2_funct_1(esk1_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
fof(c_0_12, plain, ![X9, X10]:(((r1_tarski(X9,X10)|X9!=X10)&(r1_tarski(X10,X9)|X9!=X10))&(~r1_tarski(X9,X10)|~r1_tarski(X10,X9)|X9=X10)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_13, plain, (v2_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(k5_relat_1(X2,X1))|~r1_tarski(k2_relat_1(X2),k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, negated_conjecture, (k2_relat_1(esk2_0)=k1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_18, plain, ![X11, X12, X13]:((~v2_funct_1(X11)|(~r2_hidden(X12,k1_relat_1(X11))|(~r2_hidden(X13,k1_relat_1(X11))|k1_funct_1(X11,X12)!=k1_funct_1(X11,X13))|X12=X13)|(~v1_relat_1(X11)|~v1_funct_1(X11)))&(((r2_hidden(esk3_1(X11),k1_relat_1(X11))|v2_funct_1(X11)|(~v1_relat_1(X11)|~v1_funct_1(X11)))&((r2_hidden(esk4_1(X11),k1_relat_1(X11))|v2_funct_1(X11)|(~v1_relat_1(X11)|~v1_funct_1(X11)))&(k1_funct_1(X11,esk3_1(X11))=k1_funct_1(X11,esk4_1(X11))|v2_funct_1(X11)|(~v1_relat_1(X11)|~v1_funct_1(X11)))))&(esk3_1(X11)!=esk4_1(X11)|v2_funct_1(X11)|(~v1_relat_1(X11)|~v1_funct_1(X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_funct_1])])])])])).
cnf(c_0_19, negated_conjecture, (v2_funct_1(esk2_0)|~r1_tarski(k1_relat_1(esk1_0),k1_relat_1(X1))|~v2_funct_1(k5_relat_1(esk2_0,X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16])])).
cnf(c_0_20, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_21, negated_conjecture, (v2_funct_1(k5_relat_1(esk2_0,esk1_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_22, negated_conjecture, (v1_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_23, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_24, plain, (r2_hidden(esk3_1(X1),k1_relat_1(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, negated_conjecture, (~v2_funct_1(esk2_0)|~v2_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_26, negated_conjecture, (v2_funct_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_22]), c_0_23])])).
fof(c_0_27, plain, ![X16, X17, X18, X20, X21, X22, X24]:((((r2_hidden(esk5_3(X16,X17,X18),k1_relat_1(X16))|~r2_hidden(X18,X17)|X17!=k2_relat_1(X16)|(~v1_relat_1(X16)|~v1_funct_1(X16)))&(X18=k1_funct_1(X16,esk5_3(X16,X17,X18))|~r2_hidden(X18,X17)|X17!=k2_relat_1(X16)|(~v1_relat_1(X16)|~v1_funct_1(X16))))&(~r2_hidden(X21,k1_relat_1(X16))|X20!=k1_funct_1(X16,X21)|r2_hidden(X20,X17)|X17!=k2_relat_1(X16)|(~v1_relat_1(X16)|~v1_funct_1(X16))))&((~r2_hidden(esk6_2(X16,X22),X22)|(~r2_hidden(X24,k1_relat_1(X16))|esk6_2(X16,X22)!=k1_funct_1(X16,X24))|X22=k2_relat_1(X16)|(~v1_relat_1(X16)|~v1_funct_1(X16)))&((r2_hidden(esk7_2(X16,X22),k1_relat_1(X16))|r2_hidden(esk6_2(X16,X22),X22)|X22=k2_relat_1(X16)|(~v1_relat_1(X16)|~v1_funct_1(X16)))&(esk6_2(X16,X22)=k1_funct_1(X16,esk7_2(X16,X22))|r2_hidden(esk6_2(X16,X22),X22)|X22=k2_relat_1(X16)|(~v1_relat_1(X16)|~v1_funct_1(X16)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_funct_1])])])])])])).
cnf(c_0_28, negated_conjecture, (r2_hidden(esk3_1(esk1_0),k1_relat_1(esk1_0))|v2_funct_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_22]), c_0_23])])).
cnf(c_0_29, negated_conjecture, (~v2_funct_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_26])])).
cnf(c_0_30, plain, (r2_hidden(esk4_1(X1),k1_relat_1(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_31, plain, ![X26, X27]:(~v1_relat_1(X26)|(~v1_relat_1(X27)|(~r1_tarski(k2_relat_1(X26),k1_relat_1(X27))|k1_relat_1(k5_relat_1(X26,X27))=k1_relat_1(X26)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t46_relat_1])])])).
fof(c_0_32, plain, ![X28, X29]:((v1_relat_1(k5_relat_1(X28,X29))|(~v1_relat_1(X28)|~v1_funct_1(X28)|(~v1_relat_1(X29)|~v1_funct_1(X29))))&(v1_funct_1(k5_relat_1(X28,X29))|(~v1_relat_1(X28)|~v1_funct_1(X28)|(~v1_relat_1(X29)|~v1_funct_1(X29))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc2_funct_1])])])).
fof(c_0_33, plain, ![X30, X31]:(~v1_relat_1(X30)|~v1_relat_1(X31)|v1_relat_1(k5_relat_1(X30,X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_relat_1])])).
cnf(c_0_34, plain, (r2_hidden(esk5_3(X1,X2,X3),k1_relat_1(X1))|~r2_hidden(X3,X2)|X2!=k2_relat_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35, negated_conjecture, (r2_hidden(esk3_1(esk1_0),k1_relat_1(esk1_0))), inference(sr,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_36, plain, (X1=k1_funct_1(X2,esk5_3(X2,X3,X1))|~r2_hidden(X1,X3)|X3!=k2_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, negated_conjecture, (r2_hidden(esk4_1(esk1_0),k1_relat_1(esk1_0))|v2_funct_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_22]), c_0_23])])).
cnf(c_0_38, plain, (k1_relat_1(k5_relat_1(X1,X2))=k1_relat_1(X1)|~v1_relat_1(X1)|~v1_relat_1(X2)|~r1_tarski(k2_relat_1(X1),k1_relat_1(X2))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_39, plain, (v1_funct_1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40, plain, (v1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_41, plain, ![X32, X33, X34]:(~v1_relat_1(X33)|~v1_funct_1(X33)|(~v1_relat_1(X34)|~v1_funct_1(X34)|(~r2_hidden(X32,k1_relat_1(X33))|k1_funct_1(k5_relat_1(X33,X34),X32)=k1_funct_1(X34,k1_funct_1(X33,X32))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_funct_1])])])).
cnf(c_0_42, negated_conjecture, (r2_hidden(esk5_3(X1,k1_relat_1(esk1_0),esk3_1(esk1_0)),k1_relat_1(X1))|k2_relat_1(X1)!=k1_relat_1(esk1_0)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_43, negated_conjecture, (k1_funct_1(X1,esk5_3(X1,k1_relat_1(esk1_0),esk3_1(esk1_0)))=esk3_1(esk1_0)|k2_relat_1(X1)!=k1_relat_1(esk1_0)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_36, c_0_35])).
cnf(c_0_44, negated_conjecture, (r2_hidden(esk4_1(esk1_0),k1_relat_1(esk1_0))), inference(sr,[status(thm)],[c_0_37, c_0_29])).
cnf(c_0_45, negated_conjecture, (k1_relat_1(k5_relat_1(esk2_0,X1))=k1_relat_1(esk2_0)|~r1_tarski(k1_relat_1(esk1_0),k1_relat_1(X1))|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_14]), c_0_16])])).
cnf(c_0_46, negated_conjecture, (v1_funct_1(k5_relat_1(X1,esk1_0))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_22]), c_0_23])])).
cnf(c_0_47, negated_conjecture, (v1_relat_1(k5_relat_1(X1,esk1_0))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_40, c_0_23])).
cnf(c_0_48, plain, (k1_funct_1(k5_relat_1(X1,X2),X3)=k1_funct_1(X2,k1_funct_1(X1,X3))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~r2_hidden(X3,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_49, negated_conjecture, (r2_hidden(esk5_3(esk2_0,k1_relat_1(esk1_0),esk3_1(esk1_0)),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_14]), c_0_15]), c_0_16])])).
cnf(c_0_50, negated_conjecture, (k1_funct_1(esk2_0,esk5_3(esk2_0,k1_relat_1(esk1_0),esk3_1(esk1_0)))=esk3_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_14]), c_0_15]), c_0_16])])).
cnf(c_0_51, negated_conjecture, (r2_hidden(esk5_3(X1,k1_relat_1(esk1_0),esk4_1(esk1_0)),k1_relat_1(X1))|k2_relat_1(X1)!=k1_relat_1(esk1_0)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_34, c_0_44])).
cnf(c_0_52, negated_conjecture, (k1_funct_1(X1,esk5_3(X1,k1_relat_1(esk1_0),esk4_1(esk1_0)))=esk4_1(esk1_0)|k2_relat_1(X1)!=k1_relat_1(esk1_0)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_36, c_0_44])).
cnf(c_0_53, plain, (k1_funct_1(X1,esk3_1(X1))=k1_funct_1(X1,esk4_1(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_54, plain, (X2=X3|~v2_funct_1(X1)|~r2_hidden(X2,k1_relat_1(X1))|~r2_hidden(X3,k1_relat_1(X1))|k1_funct_1(X1,X2)!=k1_funct_1(X1,X3)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_55, negated_conjecture, (k1_relat_1(k5_relat_1(esk2_0,esk1_0))=k1_relat_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_20]), c_0_23])])).
cnf(c_0_56, negated_conjecture, (v1_funct_1(k5_relat_1(esk2_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_15]), c_0_16])])).
cnf(c_0_57, negated_conjecture, (v1_relat_1(k5_relat_1(esk2_0,esk1_0))), inference(spm,[status(thm)],[c_0_47, c_0_16])).
cnf(c_0_58, negated_conjecture, (k1_funct_1(k5_relat_1(esk2_0,X1),esk5_3(esk2_0,k1_relat_1(esk1_0),esk3_1(esk1_0)))=k1_funct_1(X1,esk3_1(esk1_0))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_15]), c_0_16])]), c_0_50])).
cnf(c_0_59, negated_conjecture, (r2_hidden(esk5_3(esk2_0,k1_relat_1(esk1_0),esk4_1(esk1_0)),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_14]), c_0_15]), c_0_16])])).
cnf(c_0_60, negated_conjecture, (k1_funct_1(esk2_0,esk5_3(esk2_0,k1_relat_1(esk1_0),esk4_1(esk1_0)))=esk4_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_14]), c_0_15]), c_0_16])])).
cnf(c_0_61, negated_conjecture, (k1_funct_1(esk1_0,esk4_1(esk1_0))=k1_funct_1(esk1_0,esk3_1(esk1_0))|v2_funct_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_22]), c_0_23])])).
cnf(c_0_62, negated_conjecture, (X1=X2|k1_funct_1(k5_relat_1(esk2_0,esk1_0),X1)!=k1_funct_1(k5_relat_1(esk2_0,esk1_0),X2)|~r2_hidden(X2,k1_relat_1(esk2_0))|~r2_hidden(X1,k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_21]), c_0_56]), c_0_57])])).
cnf(c_0_63, negated_conjecture, (k1_funct_1(k5_relat_1(esk2_0,esk1_0),esk5_3(esk2_0,k1_relat_1(esk1_0),esk3_1(esk1_0)))=k1_funct_1(esk1_0,esk3_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_22]), c_0_23])])).
cnf(c_0_64, negated_conjecture, (k1_funct_1(k5_relat_1(esk2_0,X1),esk5_3(esk2_0,k1_relat_1(esk1_0),esk4_1(esk1_0)))=k1_funct_1(X1,esk4_1(esk1_0))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_59]), c_0_15]), c_0_16])]), c_0_60])).
cnf(c_0_65, negated_conjecture, (k1_funct_1(esk1_0,esk4_1(esk1_0))=k1_funct_1(esk1_0,esk3_1(esk1_0))), inference(sr,[status(thm)],[c_0_61, c_0_29])).
cnf(c_0_66, negated_conjecture, (X1=esk5_3(esk2_0,k1_relat_1(esk1_0),esk3_1(esk1_0))|k1_funct_1(k5_relat_1(esk2_0,esk1_0),X1)!=k1_funct_1(esk1_0,esk3_1(esk1_0))|~r2_hidden(X1,k1_relat_1(esk2_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_49]), c_0_63])).
cnf(c_0_67, negated_conjecture, (k1_funct_1(k5_relat_1(esk2_0,esk1_0),esk5_3(esk2_0,k1_relat_1(esk1_0),esk4_1(esk1_0)))=k1_funct_1(esk1_0,esk3_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_22]), c_0_65]), c_0_23])])).
cnf(c_0_68, plain, (v2_funct_1(X1)|esk3_1(X1)!=esk4_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_69, negated_conjecture, (esk5_3(esk2_0,k1_relat_1(esk1_0),esk4_1(esk1_0))=esk5_3(esk2_0,k1_relat_1(esk1_0),esk3_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_59]), c_0_67])])).
cnf(c_0_70, negated_conjecture, (v2_funct_1(esk1_0)|esk4_1(esk1_0)!=esk3_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_22]), c_0_23])])).
cnf(c_0_71, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_69]), c_0_50])).
cnf(c_0_72, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_71])]), c_0_29]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 73
# Proof object clause steps            : 54
# Proof object formula steps           : 19
# Proof object conjectures             : 43
# Proof object clause conjectures      : 40
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 66
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 747
# ...of these trivial                  : 3
# ...subsumed                          : 10
# ...remaining for further processing  : 734
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 127
# Generated clauses                    : 8102
# ...of the previous two non-trivial   : 8118
# Contextual simplify-reflections      : 2
# Paramodulations                      : 8076
# Factorizations                       : 0
# Equation resolutions                 : 23
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 577
#    Positive orientable unit clauses  : 446
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 130
# Current number of unprocessed clauses: 7415
# ...number of literals in the above   : 9983
# Current number of archived formulas  : 0
# Current number of archived clauses   : 155
# Clause-clause subsumption calls (NU) : 1514
# Rec. Clause-clause subsumption calls : 1207
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 3911
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 14188
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 212532

# -------------------------------------------------
# User time                : 0.062 s
# System time              : 0.003 s
# Total time               : 0.065 s
# Maximum resident set size: 28816 pages
