# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d11_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(v1_relat_1(X3)=>(X3=k7_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>(r2_hidden(X4,X2)&r2_hidden(k4_tarski(X4,X5),X1)))))), file('data/MPTP/proofs/3dc7a1fe8ad14ec9b871341ee7951496.E_input', d11_relat_1)).
fof(dt_k7_relat_1, axiom, ![X1, X2]:(v1_relat_1(X1)=>v1_relat_1(k7_relat_1(X1,X2))), file('data/MPTP/proofs/3dc7a1fe8ad14ec9b871341ee7951496.E_input', dt_k7_relat_1)).
fof(d3_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(r1_tarski(X1,X2)<=>![X3, X4]:(r2_hidden(k4_tarski(X3,X4),X1)=>r2_hidden(k4_tarski(X3,X4),X2)))), file('data/MPTP/proofs/3dc7a1fe8ad14ec9b871341ee7951496.E_input', d3_relat_1)).
fof(t106_relat_1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>![X4]:(v1_relat_1(X4)=>((r1_tarski(X3,X4)&r1_tarski(X1,X2))=>r1_tarski(k7_relat_1(X3,X1),k7_relat_1(X4,X2))))), file('data/MPTP/proofs/3dc7a1fe8ad14ec9b871341ee7951496.E_input', t106_relat_1)).
fof(t105_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(r1_tarski(X2,X3)=>r1_tarski(k7_relat_1(X2,X1),k7_relat_1(X3,X1))))), file('data/MPTP/proofs/3dc7a1fe8ad14ec9b871341ee7951496.E_input', t105_relat_1)).
fof(t102_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>k7_relat_1(k7_relat_1(X3,X1),X2)=k7_relat_1(X3,X1))), file('data/MPTP/proofs/3dc7a1fe8ad14ec9b871341ee7951496.E_input', t102_relat_1)).
fof(c_0_6, plain, ![X25, X26, X27, X28, X29, X30, X31]:((((r2_hidden(X28,X26)|~r2_hidden(k4_tarski(X28,X29),X27)|X27!=k7_relat_1(X25,X26)|~v1_relat_1(X27)|~v1_relat_1(X25))&(r2_hidden(k4_tarski(X28,X29),X25)|~r2_hidden(k4_tarski(X28,X29),X27)|X27!=k7_relat_1(X25,X26)|~v1_relat_1(X27)|~v1_relat_1(X25)))&(~r2_hidden(X30,X26)|~r2_hidden(k4_tarski(X30,X31),X25)|r2_hidden(k4_tarski(X30,X31),X27)|X27!=k7_relat_1(X25,X26)|~v1_relat_1(X27)|~v1_relat_1(X25)))&((~r2_hidden(k4_tarski(esk7_3(X25,X26,X27),esk8_3(X25,X26,X27)),X27)|(~r2_hidden(esk7_3(X25,X26,X27),X26)|~r2_hidden(k4_tarski(esk7_3(X25,X26,X27),esk8_3(X25,X26,X27)),X25))|X27=k7_relat_1(X25,X26)|~v1_relat_1(X27)|~v1_relat_1(X25))&((r2_hidden(esk7_3(X25,X26,X27),X26)|r2_hidden(k4_tarski(esk7_3(X25,X26,X27),esk8_3(X25,X26,X27)),X27)|X27=k7_relat_1(X25,X26)|~v1_relat_1(X27)|~v1_relat_1(X25))&(r2_hidden(k4_tarski(esk7_3(X25,X26,X27),esk8_3(X25,X26,X27)),X25)|r2_hidden(k4_tarski(esk7_3(X25,X26,X27),esk8_3(X25,X26,X27)),X27)|X27=k7_relat_1(X25,X26)|~v1_relat_1(X27)|~v1_relat_1(X25))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_relat_1])])])])])])).
fof(c_0_7, plain, ![X16, X17]:(~v1_relat_1(X16)|v1_relat_1(k7_relat_1(X16,X17))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_relat_1])])).
cnf(c_0_8, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(k4_tarski(X1,X2),X4)|X4!=k7_relat_1(X3,X5)|~v1_relat_1(X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, plain, (v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_10, plain, ![X18, X19, X20, X21, X22]:((~r1_tarski(X18,X19)|(~r2_hidden(k4_tarski(X20,X21),X18)|r2_hidden(k4_tarski(X20,X21),X19))|~v1_relat_1(X18))&((r2_hidden(k4_tarski(esk5_2(X18,X22),esk6_2(X18,X22)),X18)|r1_tarski(X18,X22)|~v1_relat_1(X18))&(~r2_hidden(k4_tarski(esk5_2(X18,X22),esk6_2(X18,X22)),X22)|r1_tarski(X18,X22)|~v1_relat_1(X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_relat_1])])])])])])).
fof(c_0_11, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>![X4]:(v1_relat_1(X4)=>((r1_tarski(X3,X4)&r1_tarski(X1,X2))=>r1_tarski(k7_relat_1(X3,X1),k7_relat_1(X4,X2)))))), inference(assume_negation,[status(cth)],[t106_relat_1])).
fof(c_0_12, plain, ![X10, X11, X12]:(~v1_relat_1(X11)|(~v1_relat_1(X12)|(~r1_tarski(X11,X12)|r1_tarski(k7_relat_1(X11,X10),k7_relat_1(X12,X10))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t105_relat_1])])])).
fof(c_0_13, plain, ![X13, X14, X15]:(~v1_relat_1(X15)|(~r1_tarski(X13,X14)|k7_relat_1(k7_relat_1(X15,X13),X14)=k7_relat_1(X15,X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t102_relat_1])])).
cnf(c_0_14, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(k4_tarski(X1,X2),k7_relat_1(X3,X4))|~v1_relat_1(X3)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_8]), c_0_9])).
cnf(c_0_15, plain, (r2_hidden(k4_tarski(esk5_2(X1,X2),esk6_2(X1,X2)),X1)|r1_tarski(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_16, negated_conjecture, (v1_relat_1(esk3_0)&(v1_relat_1(esk4_0)&((r1_tarski(esk3_0,esk4_0)&r1_tarski(esk1_0,esk2_0))&~r1_tarski(k7_relat_1(esk3_0,esk1_0),k7_relat_1(esk4_0,esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
cnf(c_0_17, plain, (r1_tarski(k7_relat_1(X1,X3),k7_relat_1(X2,X3))|~v1_relat_1(X1)|~v1_relat_1(X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, plain, (k7_relat_1(k7_relat_1(X1,X2),X3)=k7_relat_1(X1,X2)|~v1_relat_1(X1)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (r2_hidden(k4_tarski(X3,X4),X2)|~r1_tarski(X1,X2)|~r2_hidden(k4_tarski(X3,X4),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_20, plain, (r2_hidden(k4_tarski(esk5_2(k7_relat_1(X1,X2),X3),esk6_2(k7_relat_1(X1,X2),X3)),X1)|r1_tarski(k7_relat_1(X1,X2),X3)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_9])).
cnf(c_0_21, negated_conjecture, (~r1_tarski(k7_relat_1(esk3_0,esk1_0),k7_relat_1(esk4_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (r1_tarski(k7_relat_1(X1,X2),k7_relat_1(X3,X4))|~r1_tarski(k7_relat_1(X1,X2),X3)|~r1_tarski(X2,X4)|~v1_relat_1(X3)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_9])).
cnf(c_0_23, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, negated_conjecture, (v1_relat_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26, plain, (r1_tarski(X1,X2)|~r2_hidden(k4_tarski(esk5_2(X1,X2),esk6_2(X1,X2)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_27, plain, (r2_hidden(k4_tarski(esk5_2(k7_relat_1(X1,X2),X3),esk6_2(k7_relat_1(X1,X2),X3)),X4)|r1_tarski(k7_relat_1(X1,X2),X3)|~r1_tarski(X1,X4)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_28, negated_conjecture, (~r1_tarski(k7_relat_1(esk3_0,esk1_0),esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24]), c_0_25])])).
cnf(c_0_29, plain, (r1_tarski(k7_relat_1(X1,X2),X3)|~r1_tarski(X1,X3)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_9])).
cnf(c_0_30, negated_conjecture, (r1_tarski(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_31, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30]), c_0_25])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 32
# Proof object clause steps            : 19
# Proof object formula steps           : 13
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 7278
# ...of these trivial                  : 0
# ...subsumed                          : 5873
# ...remaining for further processing  : 1405
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 318
# Backward-rewritten                   : 0
# Generated clauses                    : 131278
# ...of the previous two non-trivial   : 125088
# Contextual simplify-reflections      : 488
# Paramodulations                      : 131275
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1067
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 1061
# Current number of unprocessed clauses: 115383
# ...number of literals in the above   : 996805
# Current number of archived formulas  : 0
# Current number of archived clauses   : 335
# Clause-clause subsumption calls (NU) : 470482
# Rec. Clause-clause subsumption calls : 70417
# Non-unit clause-clause subsumptions  : 6665
# Unit Clause-clause subsumption calls : 10
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3831434

# -------------------------------------------------
# User time                : 1.869 s
# System time              : 0.043 s
# Total time               : 1.912 s
# Maximum resident set size: 103496 pages
