# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t41_yellow_6, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m2_yellow_6(X3,X1,X2)=>r1_tarski(k10_yellow_6(X1,X2),k10_yellow_6(X1,X3))))), file('data/MPTP/proofs/3dc913cad8144837950765575818b068.E_input', t41_yellow_6)).
fof(t27_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m2_yellow_6(X3,X1,X2)=>![X4]:(r2_waybel_0(X1,X3,X4)=>r2_waybel_0(X1,X2,X4))))), file('data/MPTP/proofs/3dc913cad8144837950765575818b068.E_input', t27_yellow_6)).
fof(dt_m2_yellow_6, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&(~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1)))))=>![X3]:(m2_yellow_6(X3,X1,X2)=>(~(v2_struct_0(X3))&(v4_orders_2(X3)&(v7_waybel_0(X3)&l1_waybel_0(X3,X1)))))), file('data/MPTP/proofs/3dc913cad8144837950765575818b068.E_input', dt_m2_yellow_6)).
fof(d18_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k10_yellow_6(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>![X5]:(m1_connsp_2(X5,X1,X4)=>r1_waybel_0(X1,X2,X5)))))))), file('data/MPTP/proofs/3dc913cad8144837950765575818b068.E_input', d18_yellow_6)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/3dc913cad8144837950765575818b068.E_input', t4_subset)).
fof(t9_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r1_waybel_0(X1,X2,X3)<=>~(r2_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3)))))), file('data/MPTP/proofs/3dc913cad8144837950765575818b068.E_input', t9_waybel_0)).
fof(dt_k10_yellow_6, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&(~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1)))))=>m1_subset_1(k10_yellow_6(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/3dc913cad8144837950765575818b068.E_input', dt_k10_yellow_6)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/3dc913cad8144837950765575818b068.E_input', d3_tarski)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/3dc913cad8144837950765575818b068.E_input', dt_l1_pre_topc)).
fof(c_0_9, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m2_yellow_6(X3,X1,X2)=>r1_tarski(k10_yellow_6(X1,X2),k10_yellow_6(X1,X3)))))), inference(assume_negation,[status(cth)],[t41_yellow_6])).
fof(c_0_10, plain, ![X27, X28, X29, X30]:(v2_struct_0(X27)|~l1_struct_0(X27)|(v2_struct_0(X28)|(~v4_orders_2(X28)|(~v7_waybel_0(X28)|~l1_waybel_0(X28,X27)))|(~m2_yellow_6(X29,X27,X28)|(~r2_waybel_0(X27,X29,X30)|r2_waybel_0(X27,X28,X30))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t27_yellow_6])])])])).
fof(c_0_11, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&(v4_orders_2(esk2_0)&(v7_waybel_0(esk2_0)&l1_waybel_0(esk2_0,esk1_0))))&(m2_yellow_6(esk3_0,esk1_0,esk2_0)&~r1_tarski(k10_yellow_6(esk1_0,esk2_0),k10_yellow_6(esk1_0,esk3_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])).
fof(c_0_12, plain, ![X31, X32, X33]:((~v2_struct_0(X33)|~m2_yellow_6(X33,X31,X32)|(v2_struct_0(X31)|~l1_struct_0(X31)|(v2_struct_0(X32)|(~v4_orders_2(X32)|(~v7_waybel_0(X32)|~l1_waybel_0(X32,X31))))))&((v4_orders_2(X33)|~m2_yellow_6(X33,X31,X32)|(v2_struct_0(X31)|~l1_struct_0(X31)|(v2_struct_0(X32)|(~v4_orders_2(X32)|(~v7_waybel_0(X32)|~l1_waybel_0(X32,X31))))))&((v7_waybel_0(X33)|~m2_yellow_6(X33,X31,X32)|(v2_struct_0(X31)|~l1_struct_0(X31)|(v2_struct_0(X32)|(~v4_orders_2(X32)|(~v7_waybel_0(X32)|~l1_waybel_0(X32,X31))))))&(l1_waybel_0(X33,X31)|~m2_yellow_6(X33,X31,X32)|(v2_struct_0(X31)|~l1_struct_0(X31)|(v2_struct_0(X32)|(~v4_orders_2(X32)|(~v7_waybel_0(X32)|~l1_waybel_0(X32,X31))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_yellow_6])])])])])).
fof(c_0_13, plain, ![X18, X19, X20, X21, X22, X26]:(((~r2_hidden(X21,X20)|(~m1_connsp_2(X22,X18,X21)|r1_waybel_0(X18,X19,X22))|~m1_subset_1(X21,u1_struct_0(X18))|X20!=k10_yellow_6(X18,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X18)))|(v2_struct_0(X19)|(~v4_orders_2(X19)|(~v7_waybel_0(X19)|~l1_waybel_0(X19,X18))))|(v2_struct_0(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18))))&((m1_connsp_2(esk4_4(X18,X19,X20,X21),X18,X21)|r2_hidden(X21,X20)|~m1_subset_1(X21,u1_struct_0(X18))|X20!=k10_yellow_6(X18,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X18)))|(v2_struct_0(X19)|(~v4_orders_2(X19)|(~v7_waybel_0(X19)|~l1_waybel_0(X19,X18))))|(v2_struct_0(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18))))&(~r1_waybel_0(X18,X19,esk4_4(X18,X19,X20,X21))|r2_hidden(X21,X20)|~m1_subset_1(X21,u1_struct_0(X18))|X20!=k10_yellow_6(X18,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X18)))|(v2_struct_0(X19)|(~v4_orders_2(X19)|(~v7_waybel_0(X19)|~l1_waybel_0(X19,X18))))|(v2_struct_0(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18))))))&((m1_subset_1(esk5_3(X18,X19,X20),u1_struct_0(X18))|X20=k10_yellow_6(X18,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X18)))|(v2_struct_0(X19)|(~v4_orders_2(X19)|(~v7_waybel_0(X19)|~l1_waybel_0(X19,X18))))|(v2_struct_0(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18))))&(((m1_connsp_2(esk6_3(X18,X19,X20),X18,esk5_3(X18,X19,X20))|~r2_hidden(esk5_3(X18,X19,X20),X20)|X20=k10_yellow_6(X18,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X18)))|(v2_struct_0(X19)|(~v4_orders_2(X19)|(~v7_waybel_0(X19)|~l1_waybel_0(X19,X18))))|(v2_struct_0(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18))))&(~r1_waybel_0(X18,X19,esk6_3(X18,X19,X20))|~r2_hidden(esk5_3(X18,X19,X20),X20)|X20=k10_yellow_6(X18,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X18)))|(v2_struct_0(X19)|(~v4_orders_2(X19)|(~v7_waybel_0(X19)|~l1_waybel_0(X19,X18))))|(v2_struct_0(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18)))))&(r2_hidden(esk5_3(X18,X19,X20),X20)|(~m1_connsp_2(X26,X18,esk5_3(X18,X19,X20))|r1_waybel_0(X18,X19,X26))|X20=k10_yellow_6(X18,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X18)))|(v2_struct_0(X19)|(~v4_orders_2(X19)|(~v7_waybel_0(X19)|~l1_waybel_0(X19,X18))))|(v2_struct_0(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d18_yellow_6])])])])])])).
fof(c_0_14, plain, ![X13, X14, X15]:(~r2_hidden(X13,X14)|~m1_subset_1(X14,k1_zfmisc_1(X15))|m1_subset_1(X13,X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_15, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r2_waybel_0(X1,X2,X4)|~l1_struct_0(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~m2_yellow_6(X3,X1,X2)|~r2_waybel_0(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, negated_conjecture, (m2_yellow_6(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, negated_conjecture, (v7_waybel_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19, negated_conjecture, (v4_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_21, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_22, plain, ![X10, X11, X12]:((~r1_waybel_0(X10,X11,X12)|~r2_waybel_0(X10,X11,k6_subset_1(u1_struct_0(X10),X12))|(v2_struct_0(X11)|~l1_waybel_0(X11,X10))|(v2_struct_0(X10)|~l1_struct_0(X10)))&(r2_waybel_0(X10,X11,k6_subset_1(u1_struct_0(X10),X12))|r1_waybel_0(X10,X11,X12)|(v2_struct_0(X11)|~l1_waybel_0(X11,X10))|(v2_struct_0(X10)|~l1_struct_0(X10)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_waybel_0])])])])])).
cnf(c_0_23, plain, (l1_waybel_0(X1,X2)|v2_struct_0(X2)|v2_struct_0(X3)|~m2_yellow_6(X1,X2,X3)|~l1_struct_0(X2)|~v4_orders_2(X3)|~v7_waybel_0(X3)|~l1_waybel_0(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_24, plain, (v2_struct_0(X2)|v2_struct_0(X3)|~v2_struct_0(X1)|~m2_yellow_6(X1,X2,X3)|~l1_struct_0(X2)|~v4_orders_2(X3)|~v7_waybel_0(X3)|~l1_waybel_0(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_25, plain, (r1_waybel_0(X4,X5,X3)|v2_struct_0(X5)|v2_struct_0(X4)|~r2_hidden(X1,X2)|~m1_connsp_2(X3,X4,X1)|~m1_subset_1(X1,u1_struct_0(X4))|X2!=k10_yellow_6(X4,X5)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X4)))|~v4_orders_2(X5)|~v7_waybel_0(X5)|~l1_waybel_0(X5,X4)|~v2_pre_topc(X4)|~l1_pre_topc(X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_26, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_27, plain, ![X16, X17]:(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))|(v2_struct_0(X17)|(~v4_orders_2(X17)|(~v7_waybel_0(X17)|~l1_waybel_0(X17,X16))))|m1_subset_1(k10_yellow_6(X16,X17),k1_zfmisc_1(u1_struct_0(X16)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k10_yellow_6])])])).
cnf(c_0_28, negated_conjecture, (r2_waybel_0(esk1_0,esk2_0,X1)|~r2_waybel_0(esk1_0,esk3_0,X1)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_18]), c_0_19])]), c_0_20]), c_0_21])).
cnf(c_0_29, plain, (r2_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3))|r1_waybel_0(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, negated_conjecture, (l1_waybel_0(esk3_0,esk1_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_16]), c_0_17]), c_0_18]), c_0_19])]), c_0_20]), c_0_21])).
cnf(c_0_31, negated_conjecture, (~l1_struct_0(esk1_0)|~v2_struct_0(esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_16]), c_0_17]), c_0_18]), c_0_19])]), c_0_20]), c_0_21])).
cnf(c_0_32, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|X4!=k10_yellow_6(X1,X2)|~m1_connsp_2(X3,X1,X5)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X5,X4)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_33, plain, (v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(k10_yellow_6(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_34, plain, ![X34, X35, X36, X37, X38]:((~r1_tarski(X34,X35)|(~r2_hidden(X36,X34)|r2_hidden(X36,X35)))&((r2_hidden(esk7_2(X37,X38),X37)|r1_tarski(X37,X38))&(~r2_hidden(esk7_2(X37,X38),X38)|r1_tarski(X37,X38)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_35, plain, (v2_struct_0(X2)|v2_struct_0(X1)|~r1_waybel_0(X1,X2,X3)|~r2_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_36, negated_conjecture, (r2_waybel_0(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1))|r1_waybel_0(esk1_0,esk3_0,X1)|~l1_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_21]), c_0_30]), c_0_31])).
cnf(c_0_37, plain, (v4_orders_2(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m2_yellow_6(X1,X2,X3)|~l1_struct_0(X2)|~v4_orders_2(X3)|~v7_waybel_0(X3)|~l1_waybel_0(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_38, plain, (v7_waybel_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m2_yellow_6(X1,X2,X3)|~l1_struct_0(X2)|~v4_orders_2(X3)|~v7_waybel_0(X3)|~l1_waybel_0(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_39, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X2)|k10_yellow_6(X1,X4)!=k10_yellow_6(X1,X2)|~m1_connsp_2(X3,X1,X5)|~r2_hidden(X5,k10_yellow_6(X1,X4))|~l1_waybel_0(X2,X1)|~l1_waybel_0(X4,X1)|~v7_waybel_0(X2)|~v7_waybel_0(X4)|~v4_orders_2(X2)|~v4_orders_2(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_40, plain, (r2_hidden(esk7_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_41, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X3)|~r2_hidden(X1,k10_yellow_6(X2,X3))|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_26, c_0_33])).
cnf(c_0_42, plain, (r2_hidden(X4,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_waybel_0(X1,X2,esk4_4(X1,X2,X3,X4))|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k10_yellow_6(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_43, negated_conjecture, (r1_waybel_0(esk1_0,esk3_0,X1)|~r1_waybel_0(esk1_0,esk2_0,X1)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_17])]), c_0_20]), c_0_21])).
cnf(c_0_44, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_45, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_46, negated_conjecture, (v4_orders_2(esk3_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_16]), c_0_17]), c_0_18]), c_0_19])]), c_0_20]), c_0_21])).
cnf(c_0_47, negated_conjecture, (v7_waybel_0(esk3_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_16]), c_0_17]), c_0_18]), c_0_19])]), c_0_20]), c_0_21])).
cnf(c_0_48, plain, (r1_waybel_0(X1,X2,X3)|r1_tarski(k10_yellow_6(X1,X4),X5)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X4)|k10_yellow_6(X1,X4)!=k10_yellow_6(X1,X2)|~m1_connsp_2(X3,X1,esk7_2(k10_yellow_6(X1,X4),X5))|~l1_waybel_0(X2,X1)|~l1_waybel_0(X4,X1)|~v7_waybel_0(X2)|~v7_waybel_0(X4)|~v4_orders_2(X2)|~v4_orders_2(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_49, plain, (m1_connsp_2(esk4_4(X1,X2,X3,X4),X1,X4)|r2_hidden(X4,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k10_yellow_6(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_50, plain, (m1_subset_1(esk7_2(k10_yellow_6(X1,X2),X3),u1_struct_0(X1))|r1_tarski(k10_yellow_6(X1,X2),X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_41, c_0_40])).
cnf(c_0_51, negated_conjecture, (r2_hidden(X1,X2)|X2!=k10_yellow_6(esk1_0,esk3_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_waybel_0(esk1_0,esk2_0,esk4_4(esk1_0,esk3_0,X2,X1))|~l1_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44]), c_0_45])]), c_0_21]), c_0_46]), c_0_47]), c_0_30]), c_0_31])).
cnf(c_0_52, plain, (r2_hidden(esk7_2(k10_yellow_6(X1,X2),X3),X4)|r1_waybel_0(X1,X5,esk4_4(X1,X6,X4,esk7_2(k10_yellow_6(X1,X2),X3)))|r1_tarski(k10_yellow_6(X1,X2),X3)|v2_struct_0(X6)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X5)|k10_yellow_6(X1,X2)!=k10_yellow_6(X1,X5)|X4!=k10_yellow_6(X1,X6)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~l1_waybel_0(X5,X1)|~l1_waybel_0(X2,X1)|~l1_waybel_0(X6,X1)|~v7_waybel_0(X5)|~v7_waybel_0(X2)|~v7_waybel_0(X6)|~v4_orders_2(X5)|~v4_orders_2(X2)|~v4_orders_2(X6)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])).
cnf(c_0_53, plain, (r1_tarski(X1,X2)|~r2_hidden(esk7_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_54, negated_conjecture, (r2_hidden(esk7_2(k10_yellow_6(esk1_0,X1),X2),X3)|r1_tarski(k10_yellow_6(esk1_0,X1),X2)|v2_struct_0(X1)|k10_yellow_6(esk1_0,X1)!=k10_yellow_6(esk1_0,esk2_0)|X3!=k10_yellow_6(esk1_0,esk3_0)|~m1_subset_1(esk7_2(k10_yellow_6(esk1_0,X1),X2),u1_struct_0(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_struct_0(esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_17]), c_0_18]), c_0_19]), c_0_44]), c_0_45])]), c_0_21]), c_0_20]), c_0_46]), c_0_47]), c_0_30]), c_0_31])).
cnf(c_0_55, negated_conjecture, (r1_tarski(k10_yellow_6(esk1_0,X1),X2)|v2_struct_0(X1)|k10_yellow_6(esk1_0,X1)!=k10_yellow_6(esk1_0,esk2_0)|X2!=k10_yellow_6(esk1_0,esk3_0)|~m1_subset_1(esk7_2(k10_yellow_6(esk1_0,X1),X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_struct_0(esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_56, negated_conjecture, (r1_tarski(k10_yellow_6(esk1_0,X1),X2)|v2_struct_0(X1)|k10_yellow_6(esk1_0,X1)!=k10_yellow_6(esk1_0,esk2_0)|X2!=k10_yellow_6(esk1_0,esk3_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_struct_0(esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_50]), c_0_44]), c_0_45])]), c_0_21])).
cnf(c_0_57, negated_conjecture, (~r1_tarski(k10_yellow_6(esk1_0,esk2_0),k10_yellow_6(esk1_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_58, negated_conjecture, (r1_tarski(k10_yellow_6(esk1_0,X1),k10_yellow_6(esk1_0,esk3_0))|v2_struct_0(X1)|k10_yellow_6(esk1_0,X1)!=k10_yellow_6(esk1_0,esk2_0)|~m1_subset_1(k10_yellow_6(esk1_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_struct_0(esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(er,[status(thm)],[c_0_56])).
cnf(c_0_59, negated_conjecture, (~m1_subset_1(k10_yellow_6(esk1_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_17]), c_0_18]), c_0_19])]), c_0_20])).
fof(c_0_60, plain, ![X9]:(~l1_pre_topc(X9)|l1_struct_0(X9)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_61, negated_conjecture, (~l1_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_33]), c_0_44]), c_0_45])]), c_0_21]), c_0_46]), c_0_47]), c_0_30]), c_0_31])).
cnf(c_0_62, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_63, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_44])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 64
# Proof object clause steps            : 45
# Proof object formula steps           : 19
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 78
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 29
# Processed clauses                    : 90
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 90
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 0
# Generated clauses                    : 43
# ...of the previous two non-trivial   : 38
# Contextual simplify-reflections      : 41
# Paramodulations                      : 41
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 55
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 44
# Current number of unprocessed clauses: 6
# ...number of literals in the above   : 98
# Current number of archived formulas  : 0
# Current number of archived clauses   : 35
# Clause-clause subsumption calls (NU) : 2336
# Rec. Clause-clause subsumption calls : 144
# Non-unit clause-clause subsumptions  : 47
# Unit Clause-clause subsumption calls : 29
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5817

# -------------------------------------------------
# User time                : 0.014 s
# System time              : 0.000 s
# Total time               : 0.014 s
# Maximum resident set size: 28816 pages
