# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t1_yellow_1, axiom, ![X1]:(u1_struct_0(k2_yellow_1(X1))=X1&u1_orders_2(k2_yellow_1(X1))=k1_yellow_1(X1)), file('data/MPTP/proofs/3e0638031c6749cabeffecb789cefe33.E_input', t1_yellow_1)).
fof(redefinition_k1_yellow_1, axiom, ![X1]:k1_yellow_1(X1)=k1_wellord2(X1), file('data/MPTP/proofs/3e0638031c6749cabeffecb789cefe33.E_input', redefinition_k1_yellow_1)).
fof(d1_wellord2, axiom, ![X1, X2]:(v1_relat_1(X2)=>(X2=k1_wellord2(X1)<=>(k3_relat_1(X2)=X1&![X3, X4]:((r2_hidden(X3,X1)&r2_hidden(X4,X1))=>(r2_hidden(k4_tarski(X3,X4),X2)<=>r1_tarski(X3,X4)))))), file('data/MPTP/proofs/3e0638031c6749cabeffecb789cefe33.E_input', d1_wellord2)).
fof(dt_k1_wellord2, axiom, ![X1]:v1_relat_1(k1_wellord2(X1)), file('data/MPTP/proofs/3e0638031c6749cabeffecb789cefe33.E_input', dt_k1_wellord2)).
fof(t3_yellow_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r3_orders_2(k2_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3))))), file('data/MPTP/proofs/3e0638031c6749cabeffecb789cefe33.E_input', t3_yellow_1)).
fof(d9_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)<=>r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1)))))), file('data/MPTP/proofs/3e0638031c6749cabeffecb789cefe33.E_input', d9_orders_2)).
fof(dt_k2_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&l1_orders_2(k2_yellow_1(X1))), file('data/MPTP/proofs/3e0638031c6749cabeffecb789cefe33.E_input', dt_k2_yellow_1)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/3e0638031c6749cabeffecb789cefe33.E_input', redefinition_r3_orders_2)).
fof(fc5_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&(v3_orders_2(k2_yellow_1(X1))&(v4_orders_2(k2_yellow_1(X1))&v5_orders_2(k2_yellow_1(X1))))), file('data/MPTP/proofs/3e0638031c6749cabeffecb789cefe33.E_input', fc5_yellow_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/3e0638031c6749cabeffecb789cefe33.E_input', t2_subset)).
fof(fc6_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>(~(v2_struct_0(k2_yellow_1(X1)))&v1_orders_2(k2_yellow_1(X1)))), file('data/MPTP/proofs/3e0638031c6749cabeffecb789cefe33.E_input', fc6_yellow_1)).
fof(c_0_11, plain, ![X12]:(u1_struct_0(k2_yellow_1(X12))=X12&u1_orders_2(k2_yellow_1(X12))=k1_yellow_1(X12)), inference(variable_rename,[status(thm)],[t1_yellow_1])).
fof(c_0_12, plain, ![X11]:k1_yellow_1(X11)=k1_wellord2(X11), inference(variable_rename,[status(thm)],[redefinition_k1_yellow_1])).
fof(c_0_13, plain, ![X5, X6, X7, X8]:(((k3_relat_1(X6)=X5|X6!=k1_wellord2(X5)|~v1_relat_1(X6))&((~r2_hidden(k4_tarski(X7,X8),X6)|r1_tarski(X7,X8)|(~r2_hidden(X7,X5)|~r2_hidden(X8,X5))|X6!=k1_wellord2(X5)|~v1_relat_1(X6))&(~r1_tarski(X7,X8)|r2_hidden(k4_tarski(X7,X8),X6)|(~r2_hidden(X7,X5)|~r2_hidden(X8,X5))|X6!=k1_wellord2(X5)|~v1_relat_1(X6))))&(((r2_hidden(esk1_2(X5,X6),X5)|k3_relat_1(X6)!=X5|X6=k1_wellord2(X5)|~v1_relat_1(X6))&(r2_hidden(esk2_2(X5,X6),X5)|k3_relat_1(X6)!=X5|X6=k1_wellord2(X5)|~v1_relat_1(X6)))&((~r2_hidden(k4_tarski(esk1_2(X5,X6),esk2_2(X5,X6)),X6)|~r1_tarski(esk1_2(X5,X6),esk2_2(X5,X6))|k3_relat_1(X6)!=X5|X6=k1_wellord2(X5)|~v1_relat_1(X6))&(r2_hidden(k4_tarski(esk1_2(X5,X6),esk2_2(X5,X6)),X6)|r1_tarski(esk1_2(X5,X6),esk2_2(X5,X6))|k3_relat_1(X6)!=X5|X6=k1_wellord2(X5)|~v1_relat_1(X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_wellord2])])])])])).
cnf(c_0_14, plain, (u1_orders_2(k2_yellow_1(X1))=k1_yellow_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (k1_yellow_1(X1)=k1_wellord2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_16, plain, ![X13]:v1_relat_1(k1_wellord2(X13)), inference(variable_rename,[status(thm)],[dt_k1_wellord2])).
cnf(c_0_17, plain, (r1_tarski(X1,X2)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,X4)|~r2_hidden(X2,X4)|X3!=k1_wellord2(X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, plain, (u1_orders_2(k2_yellow_1(X1))=k1_wellord2(X1)), inference(rw,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_19, plain, (v1_relat_1(k1_wellord2(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_20, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r3_orders_2(k2_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3)))))), inference(assume_negation,[status(cth)],[t3_yellow_1])).
cnf(c_0_21, plain, (r1_tarski(X1,X2)|X3!=u1_orders_2(k2_yellow_1(X4))|~v1_relat_1(X3)|~r2_hidden(X2,X4)|~r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)), inference(rw,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_22, plain, (v1_relat_1(u1_orders_2(k2_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_19, c_0_18])).
fof(c_0_23, plain, ![X17, X18, X19]:((~r1_orders_2(X17,X18,X19)|r2_hidden(k4_tarski(X18,X19),u1_orders_2(X17))|~m1_subset_1(X19,u1_struct_0(X17))|~m1_subset_1(X18,u1_struct_0(X17))|~l1_orders_2(X17))&(~r2_hidden(k4_tarski(X18,X19),u1_orders_2(X17))|r1_orders_2(X17,X18,X19)|~m1_subset_1(X19,u1_struct_0(X17))|~m1_subset_1(X18,u1_struct_0(X17))|~l1_orders_2(X17))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_orders_2])])])])).
fof(c_0_24, plain, ![X20]:(v1_orders_2(k2_yellow_1(X20))&l1_orders_2(k2_yellow_1(X20))), inference(variable_rename,[status(thm)],[dt_k2_yellow_1])).
fof(c_0_25, negated_conjecture, (~v1_xboole_0(esk3_0)&(m1_subset_1(esk4_0,u1_struct_0(k2_yellow_1(esk3_0)))&(m1_subset_1(esk5_0,u1_struct_0(k2_yellow_1(esk3_0)))&((~r3_orders_2(k2_yellow_1(esk3_0),esk4_0,esk5_0)|~r1_tarski(esk4_0,esk5_0))&(r3_orders_2(k2_yellow_1(esk3_0),esk4_0,esk5_0)|r1_tarski(esk4_0,esk5_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_20])])])])).
cnf(c_0_26, plain, (r1_tarski(X1,X2)|~r2_hidden(k4_tarski(X1,X2),u1_orders_2(k2_yellow_1(X3)))|~r2_hidden(X2,X3)|~r2_hidden(X1,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_21]), c_0_22])])).
cnf(c_0_27, plain, (r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28, plain, (l1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29, plain, (u1_struct_0(k2_yellow_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_30, plain, ![X21, X22, X23]:((~r3_orders_2(X21,X22,X23)|r1_orders_2(X21,X22,X23)|(v2_struct_0(X21)|(~v3_orders_2(X21)|~l1_orders_2(X21))|(~m1_subset_1(X22,u1_struct_0(X21))|~m1_subset_1(X23,u1_struct_0(X21)))))&(~r1_orders_2(X21,X22,X23)|r3_orders_2(X21,X22,X23)|(v2_struct_0(X21)|(~v3_orders_2(X21)|~l1_orders_2(X21))|(~m1_subset_1(X22,u1_struct_0(X21))|~m1_subset_1(X23,u1_struct_0(X21)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
fof(c_0_31, plain, ![X24]:(v1_orders_2(k2_yellow_1(X24))&(v3_orders_2(k2_yellow_1(X24))&(v4_orders_2(k2_yellow_1(X24))&v5_orders_2(k2_yellow_1(X24))))), inference(variable_rename,[status(thm)],[fc5_yellow_1])).
fof(c_0_32, plain, ![X25, X26]:(~m1_subset_1(X25,X26)|(v1_xboole_0(X26)|r2_hidden(X25,X26))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_33, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(k2_yellow_1(esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(k2_yellow_1(esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, plain, (r1_tarski(X1,X2)|~r1_orders_2(k2_yellow_1(X3),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3)|~r2_hidden(X2,X3)|~r2_hidden(X1,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_29]), c_0_29])])).
cnf(c_0_36, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37, plain, (v3_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39, negated_conjecture, (m1_subset_1(esk5_0,esk3_0)), inference(rw,[status(thm)],[c_0_33, c_0_29])).
cnf(c_0_40, negated_conjecture, (~v1_xboole_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_41, negated_conjecture, (m1_subset_1(esk4_0,esk3_0)), inference(rw,[status(thm)],[c_0_34, c_0_29])).
cnf(c_0_42, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r1_tarski(X1,X2)|~r2_hidden(X1,X4)|~r2_hidden(X2,X4)|X3!=k1_wellord2(X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_43, plain, ![X27]:((~v2_struct_0(k2_yellow_1(X27))|v1_xboole_0(X27))&(v1_orders_2(k2_yellow_1(X27))|v1_xboole_0(X27))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc6_yellow_1])])])])).
cnf(c_0_44, plain, (v2_struct_0(k2_yellow_1(X1))|r1_tarski(X2,X3)|~r3_orders_2(k2_yellow_1(X1),X2,X3)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~r2_hidden(X3,X1)|~r2_hidden(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37]), c_0_28]), c_0_29]), c_0_29])])).
cnf(c_0_45, negated_conjecture, (r3_orders_2(k2_yellow_1(esk3_0),esk4_0,esk5_0)|r1_tarski(esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk5_0,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])).
cnf(c_0_47, negated_conjecture, (r2_hidden(esk4_0,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_41]), c_0_40])).
cnf(c_0_48, plain, (r2_hidden(k4_tarski(X1,X2),X3)|X3!=u1_orders_2(k2_yellow_1(X4))|~v1_relat_1(X3)|~r2_hidden(X2,X4)|~r2_hidden(X1,X4)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_42, c_0_18])).
cnf(c_0_49, plain, (v1_xboole_0(X1)|~v2_struct_0(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_50, negated_conjecture, (v2_struct_0(k2_yellow_1(esk3_0))|r1_tarski(esk4_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_39]), c_0_41]), c_0_46]), c_0_47])])).
cnf(c_0_51, plain, (r1_orders_2(X3,X1,X2)|~r2_hidden(k4_tarski(X1,X2),u1_orders_2(X3))|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_52, plain, (r2_hidden(k4_tarski(X1,X2),u1_orders_2(k2_yellow_1(X3)))|~r1_tarski(X1,X2)|~r2_hidden(X2,X3)|~r2_hidden(X1,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_48]), c_0_22])])).
cnf(c_0_53, negated_conjecture, (~r3_orders_2(k2_yellow_1(esk3_0),esk4_0,esk5_0)|~r1_tarski(esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_54, negated_conjecture, (r1_tarski(esk4_0,esk5_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_40])).
cnf(c_0_55, plain, (r3_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_56, plain, (r1_orders_2(k2_yellow_1(X1),X2,X3)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~r1_tarski(X2,X3)|~r2_hidden(X3,X1)|~r2_hidden(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_28]), c_0_29]), c_0_29])])).
cnf(c_0_57, negated_conjecture, (~r3_orders_2(k2_yellow_1(esk3_0),esk4_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_54])])).
cnf(c_0_58, plain, (v2_struct_0(k2_yellow_1(X1))|r3_orders_2(k2_yellow_1(X1),X2,X3)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~r1_tarski(X2,X3)|~r2_hidden(X3,X1)|~r2_hidden(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_37]), c_0_28]), c_0_29]), c_0_29])])).
cnf(c_0_59, negated_conjecture, (v2_struct_0(k2_yellow_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_39]), c_0_41]), c_0_54]), c_0_46]), c_0_47])])).
cnf(c_0_60, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_59]), c_0_40]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 61
# Proof object clause steps            : 38
# Proof object formula steps           : 23
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 47
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 27
# Processed clauses                    : 71
# ...of these trivial                  : 2
# ...subsumed                          : 0
# ...remaining for further processing  : 69
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 5
# Generated clauses                    : 28
# ...of the previous two non-trivial   : 21
# Contextual simplify-reflections      : 0
# Paramodulations                      : 21
# Factorizations                       : 0
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 32
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 16
# Current number of unprocessed clauses: 2
# ...number of literals in the above   : 14
# Current number of archived formulas  : 0
# Current number of archived clauses   : 32
# Clause-clause subsumption calls (NU) : 101
# Rec. Clause-clause subsumption calls : 29
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2933

# -------------------------------------------------
# User time                : 0.004 s
# System time              : 0.006 s
# Total time               : 0.010 s
# Maximum resident set size: 29024 pages
