# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t128_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k1_tarski(X3),X4))<=>(X1=X3&r2_hidden(X2,X4))), file('data/MPTP/proofs/3e41296b37fc420793d420d68d655452.E_input', t128_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/3e41296b37fc420793d420d68d655452.E_input', t69_enumset1)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/3e41296b37fc420793d420d68d655452.E_input', l54_zfmisc_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/3e41296b37fc420793d420d68d655452.E_input', d1_tarski)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/3e41296b37fc420793d420d68d655452.E_input', d2_tarski)).
fof(c_0_5, negated_conjecture, ~(![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k1_tarski(X3),X4))<=>(X1=X3&r2_hidden(X2,X4)))), inference(assume_negation,[status(cth)],[t128_zfmisc_1])).
fof(c_0_6, negated_conjecture, ((~r2_hidden(k4_tarski(esk1_0,esk2_0),k2_zfmisc_1(k1_tarski(esk3_0),esk4_0))|(esk1_0!=esk3_0|~r2_hidden(esk2_0,esk4_0)))&((esk1_0=esk3_0|r2_hidden(k4_tarski(esk1_0,esk2_0),k2_zfmisc_1(k1_tarski(esk3_0),esk4_0)))&(r2_hidden(esk2_0,esk4_0)|r2_hidden(k4_tarski(esk1_0,esk2_0),k2_zfmisc_1(k1_tarski(esk3_0),esk4_0))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
fof(c_0_7, plain, ![X13]:k2_tarski(X13,X13)=k1_tarski(X13), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_8, negated_conjecture, (r2_hidden(esk2_0,esk4_0)|r2_hidden(k4_tarski(esk1_0,esk2_0),k2_zfmisc_1(k1_tarski(esk3_0),esk4_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_10, plain, ![X9, X10, X11, X12]:(((r2_hidden(X9,X11)|~r2_hidden(k4_tarski(X9,X10),k2_zfmisc_1(X11,X12)))&(r2_hidden(X10,X12)|~r2_hidden(k4_tarski(X9,X10),k2_zfmisc_1(X11,X12))))&(~r2_hidden(X9,X11)|~r2_hidden(X10,X12)|r2_hidden(k4_tarski(X9,X10),k2_zfmisc_1(X11,X12)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
fof(c_0_11, plain, ![X14, X15, X16, X17, X18, X19]:(((~r2_hidden(X16,X15)|X16=X14|X15!=k1_tarski(X14))&(X17!=X14|r2_hidden(X17,X15)|X15!=k1_tarski(X14)))&((~r2_hidden(esk5_2(X18,X19),X19)|esk5_2(X18,X19)!=X18|X19=k1_tarski(X18))&(r2_hidden(esk5_2(X18,X19),X19)|esk5_2(X18,X19)=X18|X19=k1_tarski(X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_12, negated_conjecture, (~r2_hidden(k4_tarski(esk1_0,esk2_0),k2_zfmisc_1(k1_tarski(esk3_0),esk4_0))|esk1_0!=esk3_0|~r2_hidden(esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13, negated_conjecture, (r2_hidden(esk2_0,esk4_0)|r2_hidden(k4_tarski(esk1_0,esk2_0),k2_zfmisc_1(k2_tarski(esk3_0,esk3_0),esk4_0))), inference(rw,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_14, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,X2))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, negated_conjecture, (esk1_0=esk3_0|r2_hidden(k4_tarski(esk1_0,esk2_0),k2_zfmisc_1(k1_tarski(esk3_0),esk4_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_17, plain, ![X21, X22, X23, X24, X25, X26, X27, X28]:(((~r2_hidden(X24,X23)|(X24=X21|X24=X22)|X23!=k2_tarski(X21,X22))&((X25!=X21|r2_hidden(X25,X23)|X23!=k2_tarski(X21,X22))&(X25!=X22|r2_hidden(X25,X23)|X23!=k2_tarski(X21,X22))))&(((esk6_3(X26,X27,X28)!=X26|~r2_hidden(esk6_3(X26,X27,X28),X28)|X28=k2_tarski(X26,X27))&(esk6_3(X26,X27,X28)!=X27|~r2_hidden(esk6_3(X26,X27,X28),X28)|X28=k2_tarski(X26,X27)))&(r2_hidden(esk6_3(X26,X27,X28),X28)|(esk6_3(X26,X27,X28)=X26|esk6_3(X26,X27,X28)=X27)|X28=k2_tarski(X26,X27)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_18, negated_conjecture, (esk3_0!=esk1_0|~r2_hidden(esk2_0,esk4_0)|~r2_hidden(k4_tarski(esk1_0,esk2_0),k2_zfmisc_1(k2_tarski(esk3_0,esk3_0),esk4_0))), inference(rw,[status(thm)],[c_0_12, c_0_9])).
cnf(c_0_19, negated_conjecture, (r2_hidden(esk2_0,esk4_0)), inference(csr,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_20, plain, (X1=X3|X2!=k2_tarski(X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_15, c_0_9])).
cnf(c_0_21, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_22, negated_conjecture, (esk3_0=esk1_0|r2_hidden(k4_tarski(esk1_0,esk2_0),k2_zfmisc_1(k2_tarski(esk3_0,esk3_0),esk4_0))), inference(rw,[status(thm)],[c_0_16, c_0_9])).
cnf(c_0_23, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, negated_conjecture, (esk1_0!=esk3_0|~r2_hidden(k4_tarski(esk1_0,esk2_0),k2_zfmisc_1(k2_tarski(esk3_0,esk3_0),esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_19])])).
cnf(c_0_25, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_26, plain, (X1=X2|~r2_hidden(X1,k2_tarski(X2,X2))), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_27, negated_conjecture, (esk1_0=esk3_0|r2_hidden(esk1_0,k2_tarski(esk3_0,esk3_0))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X1,X3)), inference(er,[status(thm)],[c_0_23])).
cnf(c_0_29, negated_conjecture, (esk1_0!=esk3_0|~r2_hidden(esk1_0,k2_tarski(esk3_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_19])])).
cnf(c_0_30, negated_conjecture, (esk1_0=esk3_0), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_31, plain, (r2_hidden(X1,k2_tarski(X1,X2))), inference(er,[status(thm)],[c_0_28])).
cnf(c_0_32, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_30]), c_0_30]), c_0_31])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 33
# Proof object clause steps            : 22
# Proof object formula steps           : 11
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 14
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 16
# Processed clauses                    : 43
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 42
# Other redundant clauses eliminated   : 11
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 7
# Generated clauses                    : 91
# ...of the previous two non-trivial   : 74
# Contextual simplify-reflections      : 1
# Paramodulations                      : 76
# Factorizations                       : 0
# Equation resolutions                 : 15
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 17
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 13
# Current number of unprocessed clauses: 59
# ...number of literals in the above   : 255
# Current number of archived formulas  : 0
# Current number of archived clauses   : 23
# Clause-clause subsumption calls (NU) : 183
# Rec. Clause-clause subsumption calls : 124
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 18
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1964

# -------------------------------------------------
# User time                : 0.005 s
# System time              : 0.004 s
# Total time               : 0.009 s
# Maximum resident set size: 28584 pages
