# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t168_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k10_relat_1(X2,X1)=k10_relat_1(X2,k3_xboole_0(k2_relat_1(X2),X1))), file('data/MPTP/proofs/3e7e2dd9d27c4b828f05c4d293a31fba.E_input', t168_relat_1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/3e7e2dd9d27c4b828f05c4d293a31fba.E_input', t12_setfam_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/3e7e2dd9d27c4b828f05c4d293a31fba.E_input', d7_xboole_0)).
fof(t173_relat_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>(k10_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k2_relat_1(X2),X1))), file('data/MPTP/proofs/3e7e2dd9d27c4b828f05c4d293a31fba.E_input', t173_relat_1)).
fof(t55_zfmisc_1, axiom, ![X1, X2, X3]:~((r1_xboole_0(k2_tarski(X1,X2),X3)&r2_hidden(X1,X3))), file('data/MPTP/proofs/3e7e2dd9d27c4b828f05c4d293a31fba.E_input', t55_zfmisc_1)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/3e7e2dd9d27c4b828f05c4d293a31fba.E_input', t65_xboole_1)).
fof(t171_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k10_relat_1(X1,k1_xboole_0)=k1_xboole_0), file('data/MPTP/proofs/3e7e2dd9d27c4b828f05c4d293a31fba.E_input', t171_relat_1)).
fof(d14_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(X3=k10_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(k4_tarski(X4,X5),X1)&r2_hidden(X5,X2))))), file('data/MPTP/proofs/3e7e2dd9d27c4b828f05c4d293a31fba.E_input', d14_relat_1)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/3e7e2dd9d27c4b828f05c4d293a31fba.E_input', d4_xboole_0)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/3e7e2dd9d27c4b828f05c4d293a31fba.E_input', d5_relat_1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/3e7e2dd9d27c4b828f05c4d293a31fba.E_input', t4_xboole_0)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/3e7e2dd9d27c4b828f05c4d293a31fba.E_input', symmetry_r1_xboole_0)).
fof(c_0_12, plain, ![X33, X34]:(~v1_relat_1(X34)|k10_relat_1(X34,X33)=k10_relat_1(X34,k3_xboole_0(k2_relat_1(X34),X33))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t168_relat_1])])).
fof(c_0_13, plain, ![X8, X9]:k1_setfam_1(k2_tarski(X8,X9))=k3_xboole_0(X8,X9), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_14, plain, ![X57, X58]:((~r1_xboole_0(X57,X58)|k3_xboole_0(X57,X58)=k1_xboole_0)&(k3_xboole_0(X57,X58)!=k1_xboole_0|r1_xboole_0(X57,X58))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
fof(c_0_15, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>(k10_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k2_relat_1(X2),X1)))), inference(assume_negation,[status(cth)],[t173_relat_1])).
fof(c_0_16, plain, ![X21, X22, X23]:(~r1_xboole_0(k2_tarski(X21,X22),X23)|~r2_hidden(X21,X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_zfmisc_1])])).
fof(c_0_17, plain, ![X26]:r1_xboole_0(X26,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_18, plain, (k10_relat_1(X1,X2)=k10_relat_1(X1,k3_xboole_0(k2_relat_1(X1),X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_21, plain, ![X44]:(~v1_relat_1(X44)|k10_relat_1(X44,k1_xboole_0)=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t171_relat_1])])).
fof(c_0_22, negated_conjecture, (v1_relat_1(esk2_0)&((k10_relat_1(esk2_0,esk1_0)!=k1_xboole_0|~r1_xboole_0(k2_relat_1(esk2_0),esk1_0))&(k10_relat_1(esk2_0,esk1_0)=k1_xboole_0|r1_xboole_0(k2_relat_1(esk2_0),esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).
cnf(c_0_23, plain, (~r1_xboole_0(k2_tarski(X1,X2),X3)|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_25, plain, ![X45, X46, X47, X48, X50, X51, X52, X53, X55]:((((r2_hidden(k4_tarski(X48,esk8_4(X45,X46,X47,X48)),X45)|~r2_hidden(X48,X47)|X47!=k10_relat_1(X45,X46)|~v1_relat_1(X45))&(r2_hidden(esk8_4(X45,X46,X47,X48),X46)|~r2_hidden(X48,X47)|X47!=k10_relat_1(X45,X46)|~v1_relat_1(X45)))&(~r2_hidden(k4_tarski(X50,X51),X45)|~r2_hidden(X51,X46)|r2_hidden(X50,X47)|X47!=k10_relat_1(X45,X46)|~v1_relat_1(X45)))&((~r2_hidden(esk9_3(X45,X52,X53),X53)|(~r2_hidden(k4_tarski(esk9_3(X45,X52,X53),X55),X45)|~r2_hidden(X55,X52))|X53=k10_relat_1(X45,X52)|~v1_relat_1(X45))&((r2_hidden(k4_tarski(esk9_3(X45,X52,X53),esk10_3(X45,X52,X53)),X45)|r2_hidden(esk9_3(X45,X52,X53),X53)|X53=k10_relat_1(X45,X52)|~v1_relat_1(X45))&(r2_hidden(esk10_3(X45,X52,X53),X52)|r2_hidden(esk9_3(X45,X52,X53),X53)|X53=k10_relat_1(X45,X52)|~v1_relat_1(X45))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d14_relat_1])])])])])])).
cnf(c_0_26, plain, (k10_relat_1(X1,X2)=k10_relat_1(X1,k1_setfam_1(k2_tarski(k2_relat_1(X1),X2)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_27, plain, (k1_setfam_1(k2_tarski(X1,X2))=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_20, c_0_19])).
cnf(c_0_28, plain, (k10_relat_1(X1,k1_xboole_0)=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_31, plain, (r2_hidden(esk8_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k10_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, plain, (k10_relat_1(X1,k1_xboole_0)=k10_relat_1(X1,X2)|~r1_xboole_0(k2_relat_1(X1),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_33, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)=k1_xboole_0|r1_xboole_0(k2_relat_1(esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_34, negated_conjecture, (k10_relat_1(esk2_0,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35, plain, (X1!=k10_relat_1(X2,k1_xboole_0)|~r2_hidden(X3,X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_36, plain, (r2_hidden(X1,X5)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X2,X4)|X5!=k10_relat_1(X3,X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_37, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34]), c_0_29])])).
cnf(c_0_38, negated_conjecture, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_34]), c_0_29])])).
fof(c_0_39, plain, ![X35, X36, X37, X38, X39, X40, X41, X42]:((((r2_hidden(X38,X35)|~r2_hidden(X38,X37)|X37!=k3_xboole_0(X35,X36))&(r2_hidden(X38,X36)|~r2_hidden(X38,X37)|X37!=k3_xboole_0(X35,X36)))&(~r2_hidden(X39,X35)|~r2_hidden(X39,X36)|r2_hidden(X39,X37)|X37!=k3_xboole_0(X35,X36)))&((~r2_hidden(esk7_3(X40,X41,X42),X42)|(~r2_hidden(esk7_3(X40,X41,X42),X40)|~r2_hidden(esk7_3(X40,X41,X42),X41))|X42=k3_xboole_0(X40,X41))&((r2_hidden(esk7_3(X40,X41,X42),X40)|r2_hidden(esk7_3(X40,X41,X42),X42)|X42=k3_xboole_0(X40,X41))&(r2_hidden(esk7_3(X40,X41,X42),X41)|r2_hidden(esk7_3(X40,X41,X42),X42)|X42=k3_xboole_0(X40,X41))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_40, negated_conjecture, (X1!=k1_xboole_0|~r2_hidden(k4_tarski(X2,X3),esk2_0)|~r2_hidden(X3,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_29])]), c_0_38])).
fof(c_0_41, plain, ![X10, X11, X12, X14, X15, X16, X17, X19]:(((~r2_hidden(X12,X11)|r2_hidden(k4_tarski(esk3_3(X10,X11,X12),X12),X10)|X11!=k2_relat_1(X10))&(~r2_hidden(k4_tarski(X15,X14),X10)|r2_hidden(X14,X11)|X11!=k2_relat_1(X10)))&((~r2_hidden(esk4_2(X16,X17),X17)|~r2_hidden(k4_tarski(X19,esk4_2(X16,X17)),X16)|X17=k2_relat_1(X16))&(r2_hidden(esk4_2(X16,X17),X17)|r2_hidden(k4_tarski(esk5_2(X16,X17),esk4_2(X16,X17)),X16)|X17=k2_relat_1(X16)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_42, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_43, plain, ![X27, X28, X30, X31, X32]:((r1_xboole_0(X27,X28)|r2_hidden(esk6_2(X27,X28),k3_xboole_0(X27,X28)))&(~r2_hidden(X32,k3_xboole_0(X30,X31))|~r1_xboole_0(X30,X31))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
cnf(c_0_44, negated_conjecture, (~r2_hidden(k4_tarski(X1,X2),esk2_0)|~r2_hidden(X2,esk1_0)), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_45, plain, (r2_hidden(k4_tarski(esk3_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_46, plain, (r2_hidden(X1,X2)|X3!=k1_setfam_1(k2_tarski(X4,X2))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_42, c_0_19])).
cnf(c_0_47, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk6_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_48, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_49, negated_conjecture, (X1!=k2_relat_1(esk2_0)|~r2_hidden(X2,esk1_0)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_50, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_setfam_1(k2_tarski(X3,X2)))), inference(er,[status(thm)],[c_0_46])).
cnf(c_0_51, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk6_2(X1,X2),k1_setfam_1(k2_tarski(X1,X2)))), inference(rw,[status(thm)],[c_0_47, c_0_19])).
cnf(c_0_52, plain, (r2_hidden(X1,X2)|X3!=k1_setfam_1(k2_tarski(X2,X4))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_48, c_0_19])).
cnf(c_0_53, negated_conjecture, (~r2_hidden(X1,k2_relat_1(esk2_0))|~r2_hidden(X1,esk1_0)), inference(er,[status(thm)],[c_0_49])).
cnf(c_0_54, plain, (r2_hidden(esk6_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_55, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_setfam_1(k2_tarski(X2,X3)))), inference(er,[status(thm)],[c_0_52])).
fof(c_0_56, plain, ![X24, X25]:(~r1_xboole_0(X24,X25)|r1_xboole_0(X25,X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
cnf(c_0_57, negated_conjecture, (r1_xboole_0(X1,k2_relat_1(esk2_0))|~r2_hidden(esk6_2(X1,k2_relat_1(esk2_0)),esk1_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_58, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_55, c_0_51])).
cnf(c_0_59, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)!=k1_xboole_0|~r1_xboole_0(k2_relat_1(esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_60, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_61, negated_conjecture, (r1_xboole_0(esk1_0,k2_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_62, negated_conjecture, (~r1_xboole_0(k2_relat_1(esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_37])])).
cnf(c_0_63, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 64
# Proof object clause steps            : 39
# Proof object formula steps           : 25
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 16
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 28
# Processed clauses                    : 1514
# ...of these trivial                  : 12
# ...subsumed                          : 1117
# ...remaining for further processing  : 385
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 19
# Backward-rewritten                   : 32
# Generated clauses                    : 5394
# ...of the previous two non-trivial   : 4633
# Contextual simplify-reflections      : 15
# Paramodulations                      : 5245
# Factorizations                       : 22
# Equation resolutions                 : 127
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 306
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 290
# Current number of unprocessed clauses: 3118
# ...number of literals in the above   : 11831
# Current number of archived formulas  : 0
# Current number of archived clauses   : 80
# Clause-clause subsumption calls (NU) : 42248
# Rec. Clause-clause subsumption calls : 21259
# Non-unit clause-clause subsumptions  : 794
# Unit Clause-clause subsumption calls : 350
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 67065

# -------------------------------------------------
# User time                : 0.060 s
# System time              : 0.000 s
# Total time               : 0.060 s
# Maximum resident set size: 26304 pages
