# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/3f0b3aa8155e47edaa6906077638d009.E_input', t5_subset)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/3f0b3aa8155e47edaa6906077638d009.E_input', t4_subset_1)).
fof(fraenkel_a_2_1_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&(v3_orders_2(X2)&(v4_orders_2(X2)&(v5_orders_2(X2)&l1_orders_2(X2)))))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r2_hidden(X1,a_2_1_orders_2(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r2_hidden(X5,X3)=>r2_orders_2(X2,X4,X5))))))), file('data/MPTP/proofs/3f0b3aa8155e47edaa6906077638d009.E_input', fraenkel_a_2_1_orders_2)).
fof(t45_orders_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>k2_orders_2(X1,k1_struct_0(X1))=u1_struct_0(X1)), file('data/MPTP/proofs/3f0b3aa8155e47edaa6906077638d009.E_input', t45_orders_2)).
fof(d2_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k1_struct_0(X1)=k1_xboole_0), file('data/MPTP/proofs/3f0b3aa8155e47edaa6906077638d009.E_input', d2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/3f0b3aa8155e47edaa6906077638d009.E_input', dt_l1_orders_2)).
fof(t49_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(![X3]:(m1_subset_1(X3,X1)=>r2_hidden(X3,X2))=>X1=X2)), file('data/MPTP/proofs/3f0b3aa8155e47edaa6906077638d009.E_input', t49_subset_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/3f0b3aa8155e47edaa6906077638d009.E_input', fc1_xboole_0)).
fof(dt_k2_orders_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k2_orders_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/3f0b3aa8155e47edaa6906077638d009.E_input', dt_k2_orders_2)).
fof(d13_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k2_orders_2(X1,X2)=a_2_1_orders_2(X1,X2))), file('data/MPTP/proofs/3f0b3aa8155e47edaa6906077638d009.E_input', d13_orders_2)).
fof(c_0_10, plain, ![X21, X22, X23]:(~r2_hidden(X21,X22)|(~m1_subset_1(X22,k1_zfmisc_1(X23))|~v1_xboole_0(X23))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_11, plain, ![X9]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X9)), inference(variable_rename,[status(thm)],[t4_subset_1])).
fof(c_0_12, plain, ![X14, X15, X16, X18, X19]:(((m1_subset_1(esk2_3(X14,X15,X16),u1_struct_0(X15))|~r2_hidden(X14,a_2_1_orders_2(X15,X16))|(v2_struct_0(X15)|(~v3_orders_2(X15)|(~v4_orders_2(X15)|(~v5_orders_2(X15)|~l1_orders_2(X15))))|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X15)))))&((X14=esk2_3(X14,X15,X16)|~r2_hidden(X14,a_2_1_orders_2(X15,X16))|(v2_struct_0(X15)|(~v3_orders_2(X15)|(~v4_orders_2(X15)|(~v5_orders_2(X15)|~l1_orders_2(X15))))|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X15)))))&(~m1_subset_1(X18,u1_struct_0(X15))|(~r2_hidden(X18,X16)|r2_orders_2(X15,esk2_3(X14,X15,X16),X18))|~r2_hidden(X14,a_2_1_orders_2(X15,X16))|(v2_struct_0(X15)|(~v3_orders_2(X15)|(~v4_orders_2(X15)|(~v5_orders_2(X15)|~l1_orders_2(X15))))|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X15)))))))&((m1_subset_1(esk3_4(X14,X15,X16,X19),u1_struct_0(X15))|X14!=X19|~m1_subset_1(X19,u1_struct_0(X15))|r2_hidden(X14,a_2_1_orders_2(X15,X16))|(v2_struct_0(X15)|(~v3_orders_2(X15)|(~v4_orders_2(X15)|(~v5_orders_2(X15)|~l1_orders_2(X15))))|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X15)))))&((r2_hidden(esk3_4(X14,X15,X16,X19),X16)|X14!=X19|~m1_subset_1(X19,u1_struct_0(X15))|r2_hidden(X14,a_2_1_orders_2(X15,X16))|(v2_struct_0(X15)|(~v3_orders_2(X15)|(~v4_orders_2(X15)|(~v5_orders_2(X15)|~l1_orders_2(X15))))|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X15)))))&(~r2_orders_2(X15,X19,esk3_4(X14,X15,X16,X19))|X14!=X19|~m1_subset_1(X19,u1_struct_0(X15))|r2_hidden(X14,a_2_1_orders_2(X15,X16))|(v2_struct_0(X15)|(~v3_orders_2(X15)|(~v4_orders_2(X15)|(~v5_orders_2(X15)|~l1_orders_2(X15))))|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X15)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_1_orders_2])])])])])])).
cnf(c_0_13, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (r2_hidden(esk3_4(X1,X2,X3,X4),X3)|r2_hidden(X1,a_2_1_orders_2(X2,X3))|v2_struct_0(X2)|X1!=X4|~m1_subset_1(X4,u1_struct_0(X2))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_17, plain, (r2_hidden(esk3_4(X1,X2,X3,X1),X3)|r2_hidden(X1,a_2_1_orders_2(X2,X3))|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(er,[status(thm)],[c_0_15])).
fof(c_0_18, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>k2_orders_2(X1,k1_struct_0(X1))=u1_struct_0(X1))), inference(assume_negation,[status(cth)],[t45_orders_2])).
fof(c_0_19, plain, ![X12]:(~l1_struct_0(X12)|k1_struct_0(X12)=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_struct_0])])).
fof(c_0_20, plain, ![X13]:(~l1_orders_2(X13)|l1_struct_0(X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
fof(c_0_21, plain, ![X24, X25]:((m1_subset_1(esk4_2(X24,X25),X24)|X24=X25|~m1_subset_1(X25,k1_zfmisc_1(X24)))&(~r2_hidden(esk4_2(X24,X25),X25)|X24=X25|~m1_subset_1(X25,k1_zfmisc_1(X24)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t49_subset_1])])])])).
cnf(c_0_22, plain, (r2_hidden(X1,a_2_1_orders_2(X2,k1_xboole_0))|v2_struct_0(X2)|~v1_xboole_0(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_14])])).
cnf(c_0_23, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
fof(c_0_24, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&k2_orders_2(esk1_0,k1_struct_0(esk1_0))!=u1_struct_0(esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_18])])])])).
cnf(c_0_25, plain, (k1_struct_0(X1)=k1_xboole_0|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (X1=X2|~r2_hidden(esk4_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (r2_hidden(X1,a_2_1_orders_2(X2,k1_xboole_0))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
fof(c_0_29, plain, ![X10, X11]:(v2_struct_0(X10)|(~v3_orders_2(X10)|(~v4_orders_2(X10)|(~v5_orders_2(X10)|~l1_orders_2(X10))))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|m1_subset_1(k2_orders_2(X10,X11),k1_zfmisc_1(u1_struct_0(X10)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_orders_2])])])).
fof(c_0_30, plain, ![X7, X8]:(v2_struct_0(X7)|(~v3_orders_2(X7)|(~v4_orders_2(X7)|(~v5_orders_2(X7)|~l1_orders_2(X7))))|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|k2_orders_2(X7,X8)=a_2_1_orders_2(X7,X8))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_orders_2])])])])).
cnf(c_0_31, negated_conjecture, (k2_orders_2(esk1_0,k1_struct_0(esk1_0))!=u1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, plain, (k1_struct_0(X1)=k1_xboole_0|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_33, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, plain, (X1=a_2_1_orders_2(X2,k1_xboole_0)|v2_struct_0(X2)|~m1_subset_1(esk4_2(X1,a_2_1_orders_2(X2,k1_xboole_0)),u1_struct_0(X2))|~m1_subset_1(a_2_1_orders_2(X2,k1_xboole_0),k1_zfmisc_1(X1))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_35, plain, (m1_subset_1(esk4_2(X1,X2),X1)|X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_36, plain, (v2_struct_0(X1)|m1_subset_1(k2_orders_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, plain, (v2_struct_0(X1)|k2_orders_2(X1,X2)=a_2_1_orders_2(X1,X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, negated_conjecture, (k2_orders_2(esk1_0,k1_xboole_0)!=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])])).
cnf(c_0_39, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_40, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_41, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_42, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_43, plain, (a_2_1_orders_2(X1,k1_xboole_0)=u1_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(a_2_1_orders_2(X1,k1_xboole_0),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_44, plain, (m1_subset_1(a_2_1_orders_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_45, negated_conjecture, (a_2_1_orders_2(esk1_0,k1_xboole_0)!=u1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_37]), c_0_14]), c_0_33]), c_0_39]), c_0_40]), c_0_41])]), c_0_42])).
cnf(c_0_46, plain, (a_2_1_orders_2(X1,k1_xboole_0)=u1_struct_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_14])])).
cnf(c_0_47, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_33]), c_0_39]), c_0_40]), c_0_41])]), c_0_42]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 48
# Proof object clause steps            : 28
# Proof object formula steps           : 20
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 20
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 21
# Processed clauses                    : 70
# ...of these trivial                  : 0
# ...subsumed                          : 4
# ...remaining for further processing  : 66
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 0
# Generated clauses                    : 42
# ...of the previous two non-trivial   : 35
# Contextual simplify-reflections      : 1
# Paramodulations                      : 39
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 40
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 30
# Current number of unprocessed clauses: 7
# ...number of literals in the above   : 87
# Current number of archived formulas  : 0
# Current number of archived clauses   : 23
# Clause-clause subsumption calls (NU) : 1259
# Rec. Clause-clause subsumption calls : 56
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 3
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3556

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.004 s
# Total time               : 0.011 s
# Maximum resident set size: 28700 pages
