# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t63_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_xboole_0(X2,X3))=>r1_xboole_0(X1,X3)), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', t63_xboole_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', t3_subset)).
fof(t46_setfam_1, conjecture, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>~((~(X2=k1_xboole_0)&k7_setfam_1(X1,X2)=k1_xboole_0))), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', t46_setfam_1)).
fof(t79_xboole_1, axiom, ![X1, X2]:r1_xboole_0(k4_xboole_0(X1,X2),X2), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', t79_xboole_1)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', redefinition_k6_subset_1)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', symmetry_r1_xboole_0)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', t5_subset)).
fof(dt_k6_subset_1, axiom, ![X1, X2]:m1_subset_1(k6_subset_1(X1,X2),k1_zfmisc_1(X1)), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', dt_k6_subset_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', t48_xboole_1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', t12_setfam_1)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', t83_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', t28_xboole_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', t7_xboole_0)).
fof(d8_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))=>(X3=k7_setfam_1(X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(X1))=>(r2_hidden(X4,X3)<=>r2_hidden(k3_subset_1(X1,X4),X2)))))), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', d8_setfam_1)).
fof(dt_k7_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>m1_subset_1(k7_setfam_1(X1,X2),k1_zfmisc_1(k1_zfmisc_1(X1)))), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', dt_k7_setfam_1)).
fof(d5_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', d5_subset_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', commutativity_k2_tarski)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', t113_zfmisc_1)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', t4_subset)).
fof(t152_zfmisc_1, axiom, ![X1, X2]:~(r2_hidden(X1,k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', t152_zfmisc_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/3f0bcab27fa24daf9b08dbed4608d82a.E_input', t2_subset)).
fof(c_0_21, plain, ![X16, X17, X18]:(~r1_tarski(X16,X17)|~r1_xboole_0(X17,X18)|r1_xboole_0(X16,X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_xboole_1])])).
fof(c_0_22, plain, ![X19, X20]:((~m1_subset_1(X19,k1_zfmisc_1(X20))|r1_tarski(X19,X20))&(~r1_tarski(X19,X20)|m1_subset_1(X19,k1_zfmisc_1(X20)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_23, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>~((~(X2=k1_xboole_0)&k7_setfam_1(X1,X2)=k1_xboole_0)))), inference(assume_negation,[status(cth)],[t46_setfam_1])).
fof(c_0_24, plain, ![X23, X24]:r1_xboole_0(k4_xboole_0(X23,X24),X24), inference(variable_rename,[status(thm)],[t79_xboole_1])).
fof(c_0_25, plain, ![X12, X13]:k6_subset_1(X12,X13)=k4_xboole_0(X12,X13), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
cnf(c_0_26, plain, (r1_xboole_0(X1,X3)|~r1_tarski(X1,X2)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_28, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))&(esk2_0!=k1_xboole_0&k7_setfam_1(esk1_0,esk2_0)=k1_xboole_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_23])])])])).
fof(c_0_29, plain, ![X27, X28]:(~r1_xboole_0(X27,X28)|r1_xboole_0(X28,X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
cnf(c_0_30, plain, (r1_xboole_0(k4_xboole_0(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_32, plain, ![X5, X6, X7]:(~r2_hidden(X5,X6)|(~m1_subset_1(X6,k1_zfmisc_1(X7))|~v1_xboole_0(X7))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_33, plain, ![X8, X9]:m1_subset_1(k6_subset_1(X8,X9),k1_zfmisc_1(X8)), inference(variable_rename,[status(thm)],[dt_k6_subset_1])).
fof(c_0_34, plain, ![X10, X11]:k4_xboole_0(X10,k4_xboole_0(X10,X11))=k3_xboole_0(X10,X11), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_35, plain, ![X14, X15]:k1_setfam_1(k2_tarski(X14,X15))=k3_xboole_0(X14,X15), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_36, plain, ![X42, X43]:((~r1_xboole_0(X42,X43)|k4_xboole_0(X42,X43)=X42)&(k4_xboole_0(X42,X43)!=X42|r1_xboole_0(X42,X43))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
cnf(c_0_37, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_38, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_39, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_40, plain, (r1_xboole_0(k6_subset_1(X1,X2),X2)), inference(rw,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_41, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_42, plain, (m1_subset_1(k6_subset_1(X1,X2),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_43, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_44, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_45, plain, ![X25, X26]:(~r1_tarski(X25,X26)|k3_xboole_0(X25,X26)=X25), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_46, plain, (k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_47, negated_conjecture, (r1_xboole_0(esk2_0,X1)|~r1_xboole_0(k1_zfmisc_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_48, plain, (r1_xboole_0(X1,k6_subset_1(X2,X1))), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_49, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k6_subset_1(X1,X3))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_50, plain, (k6_subset_1(X1,k6_subset_1(X1,X2))=k1_setfam_1(k2_tarski(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_31]), c_0_31]), c_0_44])).
fof(c_0_51, plain, ![X44]:(X44=k1_xboole_0|r2_hidden(esk4_1(X44),X44)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
fof(c_0_52, plain, ![X29, X30, X31, X32]:(((~r2_hidden(X32,X31)|r2_hidden(k3_subset_1(X29,X32),X30)|~m1_subset_1(X32,k1_zfmisc_1(X29))|X31!=k7_setfam_1(X29,X30)|~m1_subset_1(X31,k1_zfmisc_1(k1_zfmisc_1(X29)))|~m1_subset_1(X30,k1_zfmisc_1(k1_zfmisc_1(X29))))&(~r2_hidden(k3_subset_1(X29,X32),X30)|r2_hidden(X32,X31)|~m1_subset_1(X32,k1_zfmisc_1(X29))|X31!=k7_setfam_1(X29,X30)|~m1_subset_1(X31,k1_zfmisc_1(k1_zfmisc_1(X29)))|~m1_subset_1(X30,k1_zfmisc_1(k1_zfmisc_1(X29)))))&((m1_subset_1(esk3_3(X29,X30,X31),k1_zfmisc_1(X29))|X31=k7_setfam_1(X29,X30)|~m1_subset_1(X31,k1_zfmisc_1(k1_zfmisc_1(X29)))|~m1_subset_1(X30,k1_zfmisc_1(k1_zfmisc_1(X29))))&((~r2_hidden(esk3_3(X29,X30,X31),X31)|~r2_hidden(k3_subset_1(X29,esk3_3(X29,X30,X31)),X30)|X31=k7_setfam_1(X29,X30)|~m1_subset_1(X31,k1_zfmisc_1(k1_zfmisc_1(X29)))|~m1_subset_1(X30,k1_zfmisc_1(k1_zfmisc_1(X29))))&(r2_hidden(esk3_3(X29,X30,X31),X31)|r2_hidden(k3_subset_1(X29,esk3_3(X29,X30,X31)),X30)|X31=k7_setfam_1(X29,X30)|~m1_subset_1(X31,k1_zfmisc_1(k1_zfmisc_1(X29)))|~m1_subset_1(X30,k1_zfmisc_1(k1_zfmisc_1(X29))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_setfam_1])])])])])).
fof(c_0_53, plain, ![X34, X35]:(~m1_subset_1(X35,k1_zfmisc_1(k1_zfmisc_1(X34)))|m1_subset_1(k7_setfam_1(X34,X35),k1_zfmisc_1(k1_zfmisc_1(X34)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_setfam_1])])).
fof(c_0_54, plain, ![X36, X37]:(~m1_subset_1(X37,k1_zfmisc_1(X36))|k3_subset_1(X36,X37)=k4_xboole_0(X36,X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
cnf(c_0_55, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_56, plain, ![X38, X39]:k2_tarski(X38,X39)=k2_tarski(X39,X38), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_57, plain, (k6_subset_1(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_46, c_0_31])).
cnf(c_0_58, negated_conjecture, (r1_xboole_0(esk2_0,k6_subset_1(X1,k1_zfmisc_1(esk1_0)))), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_59, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k1_setfam_1(k2_tarski(X1,X3)))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_60, plain, (X1=k1_xboole_0|r2_hidden(esk4_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_61, plain, (r2_hidden(X2,X4)|~r2_hidden(k3_subset_1(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(X1))|X4!=k7_setfam_1(X1,X3)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_62, plain, (m1_subset_1(k7_setfam_1(X2,X1),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_63, plain, (k3_subset_1(X2,X1)=k4_xboole_0(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_64, plain, (k1_setfam_1(k2_tarski(X1,X2))=X1|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_55, c_0_44])).
cnf(c_0_65, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
fof(c_0_66, plain, ![X46, X47]:((k2_zfmisc_1(X46,X47)!=k1_xboole_0|(X46=k1_xboole_0|X47=k1_xboole_0))&((X46!=k1_xboole_0|k2_zfmisc_1(X46,X47)=k1_xboole_0)&(X47!=k1_xboole_0|k2_zfmisc_1(X46,X47)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_67, plain, ![X50, X51, X52]:(~r2_hidden(X50,X51)|~m1_subset_1(X51,k1_zfmisc_1(X52))|m1_subset_1(X50,X52)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_68, negated_conjecture, (k6_subset_1(esk2_0,k6_subset_1(X1,k1_zfmisc_1(esk1_0)))=esk2_0), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_69, plain, (k1_setfam_1(k2_tarski(X1,X2))=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_70, plain, (r2_hidden(X1,k7_setfam_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(k3_subset_1(X2,X1),X3)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_61]), c_0_62])).
cnf(c_0_71, plain, (k3_subset_1(X2,X1)=k6_subset_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_63, c_0_31])).
cnf(c_0_72, plain, (k1_setfam_1(k2_tarski(X1,X2))=X2|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
fof(c_0_73, plain, ![X48, X49]:~r2_hidden(X48,k2_zfmisc_1(X48,X49)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_zfmisc_1])])).
cnf(c_0_74, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_75, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_76, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk2_0))), inference(spm,[status(thm)],[c_0_42, c_0_68])).
cnf(c_0_77, plain, (k1_xboole_0=X1|~r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_64, c_0_69])).
cnf(c_0_78, plain, (r2_hidden(X1,k7_setfam_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(k6_subset_1(X2,X1),X3)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_79, plain, (k6_subset_1(X1,k6_subset_1(X1,X2))=X2|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_50, c_0_72])).
cnf(c_0_80, plain, (~r2_hidden(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_81, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_74])).
fof(c_0_82, plain, ![X53, X54]:(~m1_subset_1(X53,X54)|(v1_xboole_0(X54)|r2_hidden(X53,X54))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_83, negated_conjecture, (m1_subset_1(X1,esk2_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_84, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_85, plain, (k1_xboole_0=X1|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_77, c_0_27])).
cnf(c_0_86, plain, (r2_hidden(k6_subset_1(X1,X2),k7_setfam_1(X1,X3))|~r1_tarski(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~r2_hidden(X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_42])])).
cnf(c_0_87, negated_conjecture, (k7_setfam_1(esk1_0,esk2_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_88, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_89, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_90, negated_conjecture, (m1_subset_1(esk4_1(esk2_0),esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_60]), c_0_84])).
cnf(c_0_91, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_38]), c_0_84])).
cnf(c_0_92, negated_conjecture, (~r1_tarski(X1,esk1_0)|~r2_hidden(X1,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_38])]), c_0_88])).
cnf(c_0_93, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_75, c_0_38])).
cnf(c_0_94, negated_conjecture, (r2_hidden(esk4_1(esk2_0),esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_91])).
cnf(c_0_95, negated_conjecture, (~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_27]), c_0_93])).
cnf(c_0_96, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_94, c_0_95]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 97
# Proof object clause steps            : 54
# Proof object formula steps           : 43
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 21
# Proof object generating inferences   : 23
# Proof object simplifying inferences  : 20
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 30
# Processed clauses                    : 15245
# ...of these trivial                  : 136
# ...subsumed                          : 13209
# ...remaining for further processing  : 1900
# Other redundant clauses eliminated   : 334
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 198
# Backward-rewritten                   : 43
# Generated clauses                    : 73431
# ...of the previous two non-trivial   : 61724
# Contextual simplify-reflections      : 28
# Paramodulations                      : 73088
# Factorizations                       : 0
# Equation resolutions                 : 334
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1616
#    Positive orientable unit clauses  : 82
#    Positive unorientable unit clauses: 5
#    Negative unit clauses             : 25
#    Non-unit-clauses                  : 1504
# Current number of unprocessed clauses: 45410
# ...number of literals in the above   : 166274
# Current number of archived formulas  : 0
# Current number of archived clauses   : 282
# Clause-clause subsumption calls (NU) : 366184
# Rec. Clause-clause subsumption calls : 185021
# Non-unit clause-clause subsumptions  : 7010
# Unit Clause-clause subsumption calls : 5121
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 154
# BW rewrite match successes           : 57
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1191268

# -------------------------------------------------
# User time                : 0.804 s
# System time              : 0.010 s
# Total time               : 0.814 s
# Maximum resident set size: 40428 pages
