# No SInE strategy applied
# Auto-Mode selected heuristic G_E___200_B02_F1_AE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t146_funct_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(X1,k1_relat_1(X2))=>r1_tarski(X1,k10_relat_1(X2,k9_relat_1(X2,X1))))), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t146_funct_1)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t1_xboole_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t36_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', d10_xboole_0)).
fof(t44_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k4_xboole_0(X1,X2),X3)=>r1_tarski(X1,k2_xboole_0(X2,X3))), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t44_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', commutativity_k2_xboole_0)).
fof(t43_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3)), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t43_xboole_1)).
fof(t167_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k10_relat_1(X2,X1),k1_relat_1(X2))), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t167_relat_1)).
fof(t71_relat_1, axiom, ![X1]:(k1_relat_1(k6_relat_1(X1))=X1&k2_relat_1(k6_relat_1(X1))=X1), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t71_relat_1)).
fof(fc3_funct_1, axiom, ![X1]:(v1_relat_1(k6_relat_1(X1))&v1_funct_1(k6_relat_1(X1))), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', fc3_funct_1)).
fof(t138_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>k10_relat_1(X3,k6_subset_1(X1,X2))=k6_subset_1(k10_relat_1(X3,X1),k10_relat_1(X3,X2))), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t138_funct_1)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', redefinition_k6_subset_1)).
fof(t3_xboole_1, axiom, ![X1]:(r1_tarski(X1,k1_xboole_0)=>X1=k1_xboole_0), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t3_xboole_1)).
fof(t181_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>k10_relat_1(k5_relat_1(X2,X3),X1)=k10_relat_1(X2,k10_relat_1(X3,X1)))), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t181_relat_1)).
fof(t77_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(k1_relat_1(X2),X1)=>k5_relat_1(k6_relat_1(X1),X2)=X2)), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t77_relat_1)).
fof(t169_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k10_relat_1(X1,k2_relat_1(X1))=k1_relat_1(X1)), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t169_relat_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t2_xboole_1)).
fof(t45_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>X2=k2_xboole_0(X1,k4_xboole_0(X2,X1))), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t45_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t3_boole)).
fof(t182_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>k1_relat_1(k5_relat_1(X1,X2))=k10_relat_1(X1,k1_relat_1(X2)))), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t182_relat_1)).
fof(t94_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k7_relat_1(X2,X1)=k5_relat_1(k6_relat_1(X1),X2)), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t94_relat_1)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t17_xboole_1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t12_setfam_1)).
fof(t139_funct_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>k10_relat_1(k7_relat_1(X3,X1),X2)=k3_xboole_0(X1,k10_relat_1(X3,X2))), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t139_funct_1)).
fof(t146_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k9_relat_1(X1,k1_relat_1(X1))=k2_relat_1(X1)), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t146_relat_1)).
fof(t162_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(r1_tarski(X2,X3)=>k9_relat_1(k7_relat_1(X1,X3),X2)=k9_relat_1(X1,X2))), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t162_relat_1)).
fof(dt_k7_relat_1, axiom, ![X1, X2]:(v1_relat_1(X1)=>v1_relat_1(k7_relat_1(X1,X2))), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', dt_k7_relat_1)).
fof(t87_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k1_relat_1(k7_relat_1(X2,X1)),X1)), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', t87_relat_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/3f1b01e51982400f8e47f6ce5b77505b.E_input', commutativity_k2_tarski)).
fof(c_0_29, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(X1,k1_relat_1(X2))=>r1_tarski(X1,k10_relat_1(X2,k9_relat_1(X2,X1)))))), inference(assume_negation,[status(cth)],[t146_funct_1])).
fof(c_0_30, plain, ![X14, X15, X16]:(~r1_tarski(X14,X15)|~r1_tarski(X15,X16)|r1_tarski(X14,X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_31, negated_conjecture, (v1_relat_1(esk2_0)&(r1_tarski(esk1_0,k1_relat_1(esk2_0))&~r1_tarski(esk1_0,k10_relat_1(esk2_0,k9_relat_1(esk2_0,esk1_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])])).
cnf(c_0_32, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_33, negated_conjecture, (r1_tarski(esk1_0,k1_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_34, negated_conjecture, (r1_tarski(X1,k1_relat_1(esk2_0))|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
fof(c_0_35, plain, ![X49, X50]:r1_tarski(k4_xboole_0(X49,X50),X49), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_36, plain, ![X60, X61]:(((r1_tarski(X60,X61)|X60!=X61)&(r1_tarski(X61,X60)|X60!=X61))&(~r1_tarski(X60,X61)|~r1_tarski(X61,X60)|X60=X61)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_37, negated_conjecture, (r1_tarski(X1,k1_relat_1(esk2_0))|~r1_tarski(X2,esk1_0)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_32, c_0_34])).
cnf(c_0_38, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_39, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_40, plain, ![X57, X58, X59]:(~r1_tarski(k4_xboole_0(X57,X58),X59)|r1_tarski(X57,k2_xboole_0(X58,X59))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t44_xboole_1])])).
cnf(c_0_41, negated_conjecture, (r1_tarski(X1,k1_relat_1(esk2_0))|~r1_tarski(X1,k4_xboole_0(esk1_0,X2))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_42, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_39])).
cnf(c_0_43, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(k4_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_44, negated_conjecture, (r1_tarski(k4_xboole_0(esk1_0,X1),k1_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
fof(c_0_45, plain, ![X23, X24]:k2_xboole_0(X23,X24)=k2_xboole_0(X24,X23), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_46, plain, ![X38, X39, X40]:(~r1_tarski(X38,k2_xboole_0(X39,X40))|r1_tarski(k4_xboole_0(X38,X39),X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
cnf(c_0_47, negated_conjecture, (r1_tarski(esk1_0,k2_xboole_0(X1,k1_relat_1(esk2_0)))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_48, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_49, plain, ![X9, X10]:(~v1_relat_1(X10)|r1_tarski(k10_relat_1(X10,X9),k1_relat_1(X10))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t167_relat_1])])).
fof(c_0_50, plain, ![X19]:(k1_relat_1(k6_relat_1(X19))=X19&k2_relat_1(k6_relat_1(X19))=X19), inference(variable_rename,[status(thm)],[t71_relat_1])).
fof(c_0_51, plain, ![X27]:(v1_relat_1(k6_relat_1(X27))&v1_funct_1(k6_relat_1(X27))), inference(variable_rename,[status(thm)],[fc3_funct_1])).
fof(c_0_52, plain, ![X34, X35, X36]:(~v1_relat_1(X36)|~v1_funct_1(X36)|k10_relat_1(X36,k6_subset_1(X34,X35))=k6_subset_1(k10_relat_1(X36,X34),k10_relat_1(X36,X35))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t138_funct_1])])).
fof(c_0_53, plain, ![X51, X52]:k6_subset_1(X51,X52)=k4_xboole_0(X51,X52), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
fof(c_0_54, plain, ![X37]:(~r1_tarski(X37,k1_xboole_0)|X37=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_xboole_1])])).
cnf(c_0_55, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_56, negated_conjecture, (r1_tarski(esk1_0,k2_xboole_0(k1_relat_1(esk2_0),X1))), inference(spm,[status(thm)],[c_0_47, c_0_48])).
fof(c_0_57, plain, ![X11, X12, X13]:(~v1_relat_1(X12)|(~v1_relat_1(X13)|k10_relat_1(k5_relat_1(X12,X13),X11)=k10_relat_1(X12,k10_relat_1(X13,X11)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t181_relat_1])])])).
fof(c_0_58, plain, ![X25, X26]:(~v1_relat_1(X26)|(~r1_tarski(k1_relat_1(X26),X25)|k5_relat_1(k6_relat_1(X25),X26)=X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t77_relat_1])])).
cnf(c_0_59, plain, (r1_tarski(k10_relat_1(X1,X2),k1_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_60, plain, (k1_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_61, plain, (v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_62, plain, (k10_relat_1(X1,k6_subset_1(X2,X3))=k6_subset_1(k10_relat_1(X1,X2),k10_relat_1(X1,X3))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_63, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_64, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_65, negated_conjecture, (r1_tarski(k4_xboole_0(esk1_0,k1_relat_1(esk2_0)),X1)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
fof(c_0_66, plain, ![X53]:(~v1_relat_1(X53)|k10_relat_1(X53,k2_relat_1(X53))=k1_relat_1(X53)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t169_relat_1])])).
cnf(c_0_67, plain, (k10_relat_1(k5_relat_1(X1,X2),X3)=k10_relat_1(X1,k10_relat_1(X2,X3))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_68, plain, (k5_relat_1(k6_relat_1(X2),X1)=X1|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_69, plain, (r1_tarski(k10_relat_1(k6_relat_1(X1),X2),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61])])).
fof(c_0_70, plain, ![X54]:r1_tarski(k1_xboole_0,X54), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_71, plain, ![X30, X31]:(~r1_tarski(X30,X31)|X31=k2_xboole_0(X30,k4_xboole_0(X31,X30))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t45_xboole_1])])).
fof(c_0_72, plain, ![X45]:k4_xboole_0(X45,k1_xboole_0)=X45, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_73, plain, (k10_relat_1(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k10_relat_1(X1,X2),k10_relat_1(X1,X3))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_63]), c_0_63])).
cnf(c_0_74, negated_conjecture, (k4_xboole_0(esk1_0,k1_relat_1(esk2_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_75, plain, (k10_relat_1(X1,k2_relat_1(X1))=k1_relat_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_76, plain, (k2_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_77, plain, (k10_relat_1(k6_relat_1(X1),k10_relat_1(X2,X3))=k10_relat_1(X2,X3)|~r1_tarski(k1_relat_1(X2),X1)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_61])])).
cnf(c_0_78, plain, (k10_relat_1(k6_relat_1(k1_xboole_0),X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_64, c_0_69])).
cnf(c_0_79, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_80, plain, (X2=k2_xboole_0(X1,k4_xboole_0(X2,X1))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_81, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_72])).
fof(c_0_82, plain, ![X43, X44]:(~v1_relat_1(X43)|(~v1_relat_1(X44)|k1_relat_1(k5_relat_1(X43,X44))=k10_relat_1(X43,k1_relat_1(X44)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t182_relat_1])])])).
fof(c_0_83, plain, ![X55, X56]:(~v1_relat_1(X56)|k7_relat_1(X56,X55)=k5_relat_1(k6_relat_1(X55),X56)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t94_relat_1])])).
cnf(c_0_84, negated_conjecture, (k4_xboole_0(k10_relat_1(X1,esk1_0),k10_relat_1(X1,k1_relat_1(esk2_0)))=k10_relat_1(X1,k1_xboole_0)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_85, plain, (k10_relat_1(k6_relat_1(X1),X1)=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_60]), c_0_61])])).
cnf(c_0_86, plain, (k10_relat_1(k6_relat_1(X1),k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_60]), c_0_79]), c_0_61])])).
cnf(c_0_87, plain, (v1_funct_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_88, plain, (k2_xboole_0(k1_xboole_0,X1)=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_79])])).
fof(c_0_89, plain, ![X28, X29]:r1_tarski(k3_xboole_0(X28,X29),X28), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_90, plain, ![X20, X21]:k1_setfam_1(k2_tarski(X20,X21))=k3_xboole_0(X20,X21), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_91, plain, ![X46, X47, X48]:(~v1_relat_1(X48)|k10_relat_1(k7_relat_1(X48,X46),X47)=k3_xboole_0(X46,k10_relat_1(X48,X47))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t139_funct_1])])).
fof(c_0_92, plain, ![X22]:(~v1_relat_1(X22)|k9_relat_1(X22,k1_relat_1(X22))=k2_relat_1(X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t146_relat_1])])).
fof(c_0_93, plain, ![X6, X7, X8]:(~v1_relat_1(X6)|(~r1_tarski(X7,X8)|k9_relat_1(k7_relat_1(X6,X8),X7)=k9_relat_1(X6,X7))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t162_relat_1])])])).
fof(c_0_94, plain, ![X17, X18]:(~v1_relat_1(X17)|v1_relat_1(k7_relat_1(X17,X18))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_relat_1])])).
fof(c_0_95, plain, ![X41, X42]:(~v1_relat_1(X42)|r1_tarski(k1_relat_1(k7_relat_1(X42,X41)),X41)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t87_relat_1])])).
cnf(c_0_96, plain, (k1_relat_1(k5_relat_1(X1,X2))=k10_relat_1(X1,k1_relat_1(X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_97, plain, (k7_relat_1(X1,X2)=k5_relat_1(k6_relat_1(X2),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_98, negated_conjecture, (k4_xboole_0(esk1_0,k10_relat_1(k6_relat_1(esk1_0),k1_relat_1(esk2_0)))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86]), c_0_87]), c_0_61])])).
cnf(c_0_99, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(spm,[status(thm)],[c_0_48, c_0_88])).
cnf(c_0_100, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_101, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_90])).
fof(c_0_102, plain, ![X32, X33]:k2_tarski(X32,X33)=k2_tarski(X33,X32), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_103, plain, (k10_relat_1(k7_relat_1(X1,X2),X3)=k3_xboole_0(X2,k10_relat_1(X1,X3))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_104, plain, (k9_relat_1(X1,k1_relat_1(X1))=k2_relat_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_105, plain, (k9_relat_1(k7_relat_1(X1,X3),X2)=k9_relat_1(X1,X2)|~v1_relat_1(X1)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_106, plain, (v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_107, plain, (r1_tarski(k1_relat_1(k7_relat_1(X1,X2)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_108, plain, (k10_relat_1(k6_relat_1(X1),k1_relat_1(X2))=k1_relat_1(k7_relat_1(X2,X1))|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_61])])).
cnf(c_0_109, negated_conjecture, (k10_relat_1(k6_relat_1(esk1_0),k1_relat_1(esk2_0))=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_98]), c_0_99]), c_0_69])])).
cnf(c_0_110, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_111, plain, (r1_tarski(k1_setfam_1(k2_tarski(X1,X2)),X1)), inference(rw,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_112, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_113, plain, (k10_relat_1(k7_relat_1(X1,X2),X3)=k1_setfam_1(k2_tarski(X2,k10_relat_1(X1,X3)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_103, c_0_101])).
cnf(c_0_114, plain, (k9_relat_1(X1,k1_relat_1(k7_relat_1(X1,X2)))=k2_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_106]), c_0_107])).
cnf(c_0_115, negated_conjecture, (k1_relat_1(k7_relat_1(esk2_0,esk1_0))=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_110])])).
cnf(c_0_116, plain, (r1_tarski(k1_setfam_1(k2_tarski(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_111, c_0_112])).
cnf(c_0_117, plain, (k1_setfam_1(k2_tarski(X1,k10_relat_1(X2,k2_relat_1(k7_relat_1(X2,X1)))))=k1_relat_1(k7_relat_1(X2,X1))|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_113]), c_0_106])).
cnf(c_0_118, negated_conjecture, (~r1_tarski(esk1_0,k10_relat_1(esk2_0,k9_relat_1(esk2_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_119, negated_conjecture, (k9_relat_1(esk2_0,esk1_0)=k2_relat_1(k7_relat_1(esk2_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_110])])).
cnf(c_0_120, plain, (r1_tarski(k1_relat_1(k7_relat_1(X1,X2)),k10_relat_1(X1,k2_relat_1(k7_relat_1(X1,X2))))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_121, negated_conjecture, (~r1_tarski(esk1_0,k10_relat_1(esk2_0,k2_relat_1(k7_relat_1(esk2_0,esk1_0))))), inference(rw,[status(thm)],[c_0_118, c_0_119])).
cnf(c_0_122, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_115]), c_0_110])]), c_0_121]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 123
# Proof object clause steps            : 64
# Proof object formula steps           : 59
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 33
# Proof object initial formulas used   : 29
# Proof object generating inferences   : 26
# Proof object simplifying inferences  : 38
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 29
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 33
# Processed clauses                    : 3681
# ...of these trivial                  : 76
# ...subsumed                          : 2903
# ...remaining for further processing  : 702
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 49
# Backward-rewritten                   : 86
# Generated clauses                    : 28468
# ...of the previous two non-trivial   : 23940
# Contextual simplify-reflections      : 20
# Paramodulations                      : 28466
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 565
#    Positive orientable unit clauses  : 157
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 403
# Current number of unprocessed clauses: 19907
# ...number of literals in the above   : 51175
# Current number of archived formulas  : 0
# Current number of archived clauses   : 137
# Clause-clause subsumption calls (NU) : 51634
# Rec. Clause-clause subsumption calls : 40034
# Non-unit clause-clause subsumptions  : 2679
# Unit Clause-clause subsumption calls : 188
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 783
# BW rewrite match successes           : 51
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 369029

# -------------------------------------------------
# User time                : 0.227 s
# System time              : 0.003 s
# Total time               : 0.230 s
# Maximum resident set size: 28620 pages
