# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d16_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:k11_relat_1(X1,X2)=k9_relat_1(X1,k1_tarski(X2))), file('data/MPTP/proofs/3f4d8033a37f48c297c47eb06584f49b.E_input', d16_relat_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/3f4d8033a37f48c297c47eb06584f49b.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/3f4d8033a37f48c297c47eb06584f49b.E_input', t70_enumset1)).
fof(t205_relat_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>(r2_hidden(X1,k1_relat_1(X2))<=>~(k11_relat_1(X2,X1)=k1_xboole_0))), file('data/MPTP/proofs/3f4d8033a37f48c297c47eb06584f49b.E_input', t205_relat_1)).
fof(t204_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)<=>r2_hidden(X2,k11_relat_1(X3,X1)))), file('data/MPTP/proofs/3f4d8033a37f48c297c47eb06584f49b.E_input', t204_relat_1)).
fof(d1_relat_1, axiom, ![X1]:(v1_relat_1(X1)<=>![X2]:~((r2_hidden(X2,X1)&![X3, X4]:~(X2=k4_tarski(X3,X4))))), file('data/MPTP/proofs/3f4d8033a37f48c297c47eb06584f49b.E_input', d1_relat_1)).
fof(t56_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>(![X2, X3]:~(r2_hidden(k4_tarski(X2,X3),X1))=>X1=k1_xboole_0)), file('data/MPTP/proofs/3f4d8033a37f48c297c47eb06584f49b.E_input', t56_relat_1)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/3f4d8033a37f48c297c47eb06584f49b.E_input', d4_relat_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/3f4d8033a37f48c297c47eb06584f49b.E_input', t3_xboole_0)).
fof(t151_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(k9_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k1_relat_1(X2),X1))), file('data/MPTP/proofs/3f4d8033a37f48c297c47eb06584f49b.E_input', t151_relat_1)).
fof(d1_enumset1, axiom, ![X1, X2, X3, X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3)))))), file('data/MPTP/proofs/3f4d8033a37f48c297c47eb06584f49b.E_input', d1_enumset1)).
fof(c_0_11, plain, ![X41, X42]:(~v1_relat_1(X41)|k11_relat_1(X41,X42)=k9_relat_1(X41,k1_tarski(X42))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d16_relat_1])])])).
fof(c_0_12, plain, ![X8]:k2_tarski(X8,X8)=k1_tarski(X8), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_13, plain, ![X22, X23]:k1_enumset1(X22,X22,X23)=k2_tarski(X22,X23), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_14, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>(r2_hidden(X1,k1_relat_1(X2))<=>~(k11_relat_1(X2,X1)=k1_xboole_0)))), inference(assume_negation,[status(cth)],[t205_relat_1])).
fof(c_0_15, plain, ![X24, X25, X26]:((~r2_hidden(k4_tarski(X24,X25),X26)|r2_hidden(X25,k11_relat_1(X26,X24))|~v1_relat_1(X26))&(~r2_hidden(X25,k11_relat_1(X26,X24))|r2_hidden(k4_tarski(X24,X25),X26)|~v1_relat_1(X26))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t204_relat_1])])])).
fof(c_0_16, plain, ![X43, X44, X47, X49, X50]:((~v1_relat_1(X43)|(~r2_hidden(X44,X43)|X44=k4_tarski(esk9_2(X43,X44),esk10_2(X43,X44))))&((r2_hidden(esk11_1(X47),X47)|v1_relat_1(X47))&(esk11_1(X47)!=k4_tarski(X49,X50)|v1_relat_1(X47)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_relat_1])])])])])])])).
cnf(c_0_17, plain, (k11_relat_1(X1,X2)=k9_relat_1(X1,k1_tarski(X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_20, negated_conjecture, (v1_relat_1(esk2_0)&((~r2_hidden(esk1_0,k1_relat_1(esk2_0))|k11_relat_1(esk2_0,esk1_0)=k1_xboole_0)&(r2_hidden(esk1_0,k1_relat_1(esk2_0))|k11_relat_1(esk2_0,esk1_0)!=k1_xboole_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])).
fof(c_0_21, plain, ![X27]:(~v1_relat_1(X27)|(r2_hidden(k4_tarski(esk4_1(X27),esk5_1(X27)),X27)|X27=k1_xboole_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t56_relat_1])])])])).
fof(c_0_22, plain, ![X30, X31, X32, X34, X35, X36, X37, X39]:(((~r2_hidden(X32,X31)|r2_hidden(k4_tarski(X32,esk6_3(X30,X31,X32)),X30)|X31!=k1_relat_1(X30))&(~r2_hidden(k4_tarski(X34,X35),X30)|r2_hidden(X34,X31)|X31!=k1_relat_1(X30)))&((~r2_hidden(esk7_2(X36,X37),X37)|~r2_hidden(k4_tarski(esk7_2(X36,X37),X39),X36)|X37=k1_relat_1(X36))&(r2_hidden(esk7_2(X36,X37),X37)|r2_hidden(k4_tarski(esk7_2(X36,X37),esk8_2(X36,X37)),X36)|X37=k1_relat_1(X36)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
cnf(c_0_23, plain, (r2_hidden(k4_tarski(X3,X1),X2)|~r2_hidden(X1,k11_relat_1(X2,X3))|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, plain, (r2_hidden(esk11_1(X1),X1)|v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_25, plain, ![X51, X52, X54, X55, X56]:(((r2_hidden(esk12_2(X51,X52),X51)|r1_xboole_0(X51,X52))&(r2_hidden(esk12_2(X51,X52),X52)|r1_xboole_0(X51,X52)))&(~r2_hidden(X56,X54)|~r2_hidden(X56,X55)|~r1_xboole_0(X54,X55))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_26, plain, ![X20, X21]:((k9_relat_1(X21,X20)!=k1_xboole_0|r1_xboole_0(k1_relat_1(X21),X20)|~v1_relat_1(X21))&(~r1_xboole_0(k1_relat_1(X21),X20)|k9_relat_1(X21,X20)=k1_xboole_0|~v1_relat_1(X21))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t151_relat_1])])])).
cnf(c_0_27, plain, (k11_relat_1(X1,X2)=k9_relat_1(X1,k1_enumset1(X2,X2,X2))|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_18]), c_0_19])).
cnf(c_0_28, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_29, plain, ![X9, X10, X11, X12, X13, X14, X15, X16, X17, X18]:(((~r2_hidden(X13,X12)|(X13=X9|(X13=X10|X13=X11))|X12!=k1_enumset1(X9,X10,X11))&((X14!=X9|r2_hidden(X14,X12)|X12!=k1_enumset1(X9,X10,X11))&((X14!=X10|r2_hidden(X14,X12)|X12!=k1_enumset1(X9,X10,X11))&(X14!=X11|r2_hidden(X14,X12)|X12!=k1_enumset1(X9,X10,X11)))))&(((esk3_4(X15,X16,X17,X18)!=X15|~r2_hidden(esk3_4(X15,X16,X17,X18),X18)|X18=k1_enumset1(X15,X16,X17))&((esk3_4(X15,X16,X17,X18)!=X16|~r2_hidden(esk3_4(X15,X16,X17,X18),X18)|X18=k1_enumset1(X15,X16,X17))&(esk3_4(X15,X16,X17,X18)!=X17|~r2_hidden(esk3_4(X15,X16,X17,X18),X18)|X18=k1_enumset1(X15,X16,X17))))&(r2_hidden(esk3_4(X15,X16,X17,X18),X18)|(esk3_4(X15,X16,X17,X18)=X15|(esk3_4(X15,X16,X17,X18)=X16|esk3_4(X15,X16,X17,X18)=X17))|X18=k1_enumset1(X15,X16,X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])).
cnf(c_0_30, plain, (r2_hidden(k4_tarski(esk4_1(X1),esk5_1(X1)),X1)|X1=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, plain, (r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32, plain, (r2_hidden(k4_tarski(X1,esk11_1(k11_relat_1(X2,X1))),X2)|v1_relat_1(k11_relat_1(X2,X1))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_33, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, plain, (r1_xboole_0(k1_relat_1(X1),X2)|k9_relat_1(X1,X2)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(esk2_0))|k11_relat_1(esk2_0,esk1_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_36, negated_conjecture, (k11_relat_1(esk2_0,X1)=k9_relat_1(esk2_0,k1_enumset1(X1,X1,X1))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_37, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X2,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, plain, (k11_relat_1(X1,X2)=k1_xboole_0|r2_hidden(k4_tarski(X2,k4_tarski(esk4_1(k11_relat_1(X1,X2)),esk5_1(k11_relat_1(X1,X2)))),X1)|~v1_relat_1(k11_relat_1(X1,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_23, c_0_30])).
cnf(c_0_39, plain, (r2_hidden(X1,X2)|v1_relat_1(k11_relat_1(X3,X1))|X2!=k1_relat_1(X3)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_40, plain, (k9_relat_1(X1,X2)!=k1_xboole_0|~r2_hidden(X3,k1_relat_1(X1))|~r2_hidden(X3,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_41, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(esk2_0))|k9_relat_1(esk2_0,k1_enumset1(esk1_0,esk1_0,esk1_0))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_42, plain, (r2_hidden(X1,X2)|X2!=k1_enumset1(X1,X3,X4)), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_43, negated_conjecture, (k11_relat_1(esk2_0,esk1_0)=k1_xboole_0|~r2_hidden(esk1_0,k1_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_44, plain, (k11_relat_1(X1,X2)=k1_xboole_0|r2_hidden(X2,X3)|X3!=k1_relat_1(X1)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_38]), c_0_39])).
cnf(c_0_45, negated_conjecture, (k9_relat_1(esk2_0,k1_enumset1(esk1_0,esk1_0,esk1_0))!=k1_xboole_0|k9_relat_1(esk2_0,X1)!=k1_xboole_0|~r2_hidden(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_28])])).
cnf(c_0_46, plain, (r2_hidden(X1,k1_enumset1(X1,X2,X3))), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_47, negated_conjecture, (k9_relat_1(esk2_0,k1_enumset1(esk1_0,esk1_0,esk1_0))=k1_xboole_0|~r2_hidden(esk1_0,k1_relat_1(esk2_0))), inference(rw,[status(thm)],[c_0_43, c_0_36])).
cnf(c_0_48, plain, (k11_relat_1(X1,X2)=k1_xboole_0|r2_hidden(X2,k1_relat_1(X1))|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_49, negated_conjecture, (k9_relat_1(esk2_0,k1_enumset1(esk1_0,esk1_0,esk1_0))!=k1_xboole_0|k9_relat_1(esk2_0,k1_enumset1(esk1_0,X1,X2))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_50, negated_conjecture, (k9_relat_1(esk2_0,k1_enumset1(esk1_0,esk1_0,esk1_0))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_36]), c_0_28])])).
cnf(c_0_51, negated_conjecture, (k9_relat_1(esk2_0,k1_enumset1(esk1_0,X1,X2))!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_50])])).
cnf(c_0_52, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_50, c_0_51]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 53
# Proof object clause steps            : 30
# Proof object formula steps           : 23
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 14
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 27
# Processed clauses                    : 126
# ...of these trivial                  : 1
# ...subsumed                          : 10
# ...remaining for further processing  : 115
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 9
# Generated clauses                    : 291
# ...of the previous two non-trivial   : 289
# Contextual simplify-reflections      : 1
# Paramodulations                      : 279
# Factorizations                       : 0
# Equation resolutions                 : 11
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 73
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 64
# Current number of unprocessed clauses: 208
# ...number of literals in the above   : 738
# Current number of archived formulas  : 0
# Current number of archived clauses   : 41
# Clause-clause subsumption calls (NU) : 335
# Rec. Clause-clause subsumption calls : 216
# Non-unit clause-clause subsumptions  : 13
# Unit Clause-clause subsumption calls : 82
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6442

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.000 s
# Total time               : 0.013 s
# Maximum resident set size: 28644 pages
