# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t159_funct_2, conjecture, ![X1, X2, X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>r1_tarski(k5_partfun1(X1,X2,X3),k1_funct_2(X1,X2))), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', t159_funct_2)).
fof(t14_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))=>(r1_tarski(k2_relat_1(X4),X2)=>m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', t14_relset_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', d10_xboole_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', d1_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', d3_tarski)).
fof(cc3_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', cc3_relset_1)).
fof(rc1_subset_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>?[X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))&~(v1_xboole_0(X2)))), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', rc1_subset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', t3_subset)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', cc2_relset_1)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', t1_xboole_1)).
fof(d19_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(v5_relat_1(X2,X1)<=>r1_tarski(k2_relat_1(X2),X1))), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', d19_relat_1)).
fof(cc4_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', cc4_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', cc1_relset_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', t113_zfmisc_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', fc1_xboole_0)).
fof(fc2_partfun1, axiom, ![X1, X2, X3]:((~(v1_xboole_0(X1))&(v1_xboole_0(X2)&(v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))=>v1_xboole_0(k5_partfun1(X1,X2,X3))), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', fc2_partfun1)).
fof(t11_funct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>((X2=k1_xboole_0=>X1=k1_xboole_0)=>r2_hidden(X3,k1_funct_2(X1,X2)))), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', t11_funct_2)).
fof(t158_funct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>![X4]:(r2_hidden(X4,k5_partfun1(X1,X2,X3))=>(v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', t158_funct_2)).
fof(t135_zfmisc_1, axiom, ![X1, X2]:((r1_tarski(X1,k2_zfmisc_1(X1,X2))|r1_tarski(X1,k2_zfmisc_1(X2,X1)))=>X1=k1_xboole_0), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', t135_zfmisc_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', t2_xboole_1)).
fof(t12_funct_2, axiom, ![X1, X2]:((v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>r2_hidden(X2,k1_funct_2(X1,X1))), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', t12_funct_2)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/410be16af98246dbad35e174fc2fc7eb.E_input', t4_subset_1)).
fof(c_0_22, negated_conjecture, ~(![X1, X2, X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>r1_tarski(k5_partfun1(X1,X2,X3),k1_funct_2(X1,X2)))), inference(assume_negation,[status(cth)],[t159_funct_2])).
fof(c_0_23, plain, ![X8, X9, X10, X11]:(~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X10,X8)))|(~r1_tarski(k2_relat_1(X11),X9)|m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X10,X9))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t14_relset_1])])).
fof(c_0_24, negated_conjecture, ((v1_funct_1(esk3_0)&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0))))&~r1_tarski(k5_partfun1(esk1_0,esk2_0,esk3_0),k1_funct_2(esk1_0,esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])).
fof(c_0_25, plain, ![X12, X13]:(((r1_tarski(X12,X13)|X12!=X13)&(r1_tarski(X13,X12)|X12!=X13))&(~r1_tarski(X12,X13)|~r1_tarski(X13,X12)|X12=X13)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_26, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r1_tarski(k2_relat_1(X1),X4)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_28, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_29, plain, ![X14, X15, X16]:((~v1_xboole_0(X14)|~r2_hidden(X15,X14))&(r2_hidden(esk4_1(X16),X16)|v1_xboole_0(X16))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_30, plain, ![X18, X19, X20, X21, X22]:((~r1_tarski(X18,X19)|(~r2_hidden(X20,X18)|r2_hidden(X20,X19)))&((r2_hidden(esk5_2(X21,X22),X21)|r1_tarski(X21,X22))&(~r2_hidden(esk5_2(X21,X22),X22)|r1_tarski(X21,X22)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_31, plain, ![X24, X25, X26]:(~v1_xboole_0(X24)|(~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(X24,X25)))|v1_xboole_0(X26))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_relset_1])])])).
fof(c_0_32, plain, ![X27]:((m1_subset_1(esk6_1(X27),k1_zfmisc_1(X27))|v1_xboole_0(X27))&(~v1_xboole_0(esk6_1(X27))|v1_xboole_0(X27))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[rc1_subset_1])])])])])).
fof(c_0_33, plain, ![X29, X30]:((~m1_subset_1(X29,k1_zfmisc_1(X30))|r1_tarski(X29,X30))&(~r1_tarski(X29,X30)|m1_subset_1(X29,k1_zfmisc_1(X30)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_34, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,X1)))|~r1_tarski(k2_relat_1(esk3_0),X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_35, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_28])).
fof(c_0_36, plain, ![X31, X32, X33]:((v4_relat_1(X33,X31)|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X32))))&(v5_relat_1(X33,X32)|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X32))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_37, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_40, plain, (m1_subset_1(esk6_1(X1),k1_zfmisc_1(X1))|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_41, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk6_1(X1))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_42, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_43, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0))))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
fof(c_0_44, plain, ![X34, X35, X36]:(~r1_tarski(X34,X35)|~r1_tarski(X35,X36)|r1_tarski(X34,X36)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_45, plain, ![X37, X38]:((~v5_relat_1(X38,X37)|r1_tarski(k2_relat_1(X38),X37)|~v1_relat_1(X38))&(~r1_tarski(k2_relat_1(X38),X37)|v5_relat_1(X38,X37)|~v1_relat_1(X38))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d19_relat_1])])])).
fof(c_0_46, plain, ![X39, X40, X41]:(~v1_xboole_0(X39)|(~m1_subset_1(X41,k1_zfmisc_1(k2_zfmisc_1(X40,X39)))|v1_xboole_0(X41))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])).
cnf(c_0_47, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_48, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_49, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_50, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])).
cnf(c_0_51, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_52, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0)))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_53, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_54, plain, (r1_tarski(k2_relat_1(X1),X2)|~v5_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_55, plain, ![X42, X43, X44]:(~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(X42,X43)))|v1_relat_1(X44)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_56, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_57, plain, (v5_relat_1(X1,X2)|~r1_tarski(X1,k2_zfmisc_1(X3,X2))), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_58, plain, (r1_tarski(k2_zfmisc_1(X1,X2),X3)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_59, negated_conjecture, (k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0))=esk3_0|~r1_tarski(k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0)),esk3_0)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_60, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_42, c_0_27])).
cnf(c_0_61, plain, (r1_tarski(X1,X2)|~v1_relat_1(X3)|~v5_relat_1(X3,X2)|~r1_tarski(X1,k2_relat_1(X3))), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_62, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_63, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~r1_tarski(X1,k2_zfmisc_1(X3,X2))), inference(spm,[status(thm)],[c_0_56, c_0_48])).
cnf(c_0_64, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X3)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_65, negated_conjecture, (k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0))=esk3_0|~v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_59, c_0_58])).
cnf(c_0_66, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_40]), c_0_41])).
cnf(c_0_67, plain, (v5_relat_1(X1,X2)|~r1_tarski(k2_relat_1(X1),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_68, negated_conjecture, (r1_tarski(X1,k2_zfmisc_1(esk1_0,esk2_0))|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_53, c_0_60])).
cnf(c_0_69, plain, (r1_tarski(k2_relat_1(X1),X2)|~v1_relat_1(X3)|~v1_relat_1(X1)|~v5_relat_1(X1,k2_relat_1(X3))|~v5_relat_1(X3,X2)), inference(spm,[status(thm)],[c_0_61, c_0_54])).
cnf(c_0_70, negated_conjecture, (v5_relat_1(esk3_0,k2_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_47, c_0_43])).
cnf(c_0_71, negated_conjecture, (v1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_62, c_0_27])).
fof(c_0_72, plain, ![X45, X46]:((k2_zfmisc_1(X45,X46)!=k1_xboole_0|(X45=k1_xboole_0|X46=k1_xboole_0))&((X45!=k1_xboole_0|k2_zfmisc_1(X45,X46)=k1_xboole_0)&(X46!=k1_xboole_0|k2_zfmisc_1(X45,X46)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_73, plain, (v1_xboole_0(k2_relat_1(X1))|~v1_relat_1(X1)|~v5_relat_1(X1,k2_zfmisc_1(X2,X3))|~v1_xboole_0(X3)), inference(spm,[status(thm)],[c_0_63, c_0_54])).
cnf(c_0_74, negated_conjecture, (v5_relat_1(esk3_0,X1)|~v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_75, negated_conjecture, (r1_tarski(X1,k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0)))|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_53, c_0_52])).
cnf(c_0_76, plain, (X1=k2_relat_1(X2)|~v1_relat_1(X2)|~v5_relat_1(X2,X1)|~r1_tarski(X1,k2_relat_1(X2))), inference(spm,[status(thm)],[c_0_51, c_0_54])).
cnf(c_0_77, plain, (r1_tarski(k2_zfmisc_1(X1,X2),X3)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_49, c_0_66])).
cnf(c_0_78, negated_conjecture, (v5_relat_1(X1,k2_zfmisc_1(esk1_0,esk2_0))|~v1_relat_1(X1)|~r1_tarski(k2_relat_1(X1),esk3_0)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_79, negated_conjecture, (r1_tarski(k2_relat_1(esk3_0),X1)|~v5_relat_1(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_71])])).
cnf(c_0_80, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_81, negated_conjecture, (v1_xboole_0(k2_relat_1(esk3_0))|~v1_xboole_0(esk1_0)|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_71])])).
cnf(c_0_82, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
fof(c_0_83, plain, ![X47, X48, X49]:(v1_xboole_0(X47)|(~v1_xboole_0(X48)|(~v1_funct_1(X49)|~m1_subset_1(X49,k1_zfmisc_1(k2_zfmisc_1(X47,X48)))))|v1_xboole_0(k5_partfun1(X47,X48,X49))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_partfun1])])])).
cnf(c_0_84, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0)))))|~r1_tarski(k2_relat_1(esk3_0),esk3_0)), inference(spm,[status(thm)],[c_0_34, c_0_75])).
cnf(c_0_85, plain, (k2_zfmisc_1(X1,X2)=k2_relat_1(X3)|~v1_relat_1(X3)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_86, negated_conjecture, (v5_relat_1(esk3_0,k2_zfmisc_1(esk1_0,esk2_0))|~v5_relat_1(esk3_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_71])])).
cnf(c_0_87, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_80])).
cnf(c_0_88, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X3)|~r1_tarski(X1,k2_zfmisc_1(X4,X3))), inference(spm,[status(thm)],[c_0_53, c_0_77])).
cnf(c_0_89, negated_conjecture, (k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0))=X1|~r1_tarski(k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0)),X1)|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_51, c_0_75])).
cnf(c_0_90, negated_conjecture, (v1_xboole_0(k2_relat_1(esk3_0))|~v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_91, plain, (v1_xboole_0(X1)|v1_xboole_0(k5_partfun1(X1,X2,X3))|~v1_xboole_0(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_92, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_72])).
fof(c_0_93, plain, ![X50, X51, X52]:((X51=k1_xboole_0|r2_hidden(X52,k1_funct_2(X50,X51))|(~v1_funct_1(X52)|(~v1_funct_2(X52,X50,X51)|~m1_subset_1(X52,k1_zfmisc_1(k2_zfmisc_1(X50,X51))))))&(X50!=k1_xboole_0|r2_hidden(X52,k1_funct_2(X50,X51))|(~v1_funct_1(X52)|(~v1_funct_2(X52,X50,X51)|~m1_subset_1(X52,k1_zfmisc_1(k2_zfmisc_1(X50,X51))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_funct_2])])])).
fof(c_0_94, plain, ![X53, X54, X55, X56]:((v1_funct_1(X56)|~r2_hidden(X56,k5_partfun1(X53,X54,X55))|(~v1_funct_1(X55)|~m1_subset_1(X55,k1_zfmisc_1(k2_zfmisc_1(X53,X54)))))&((v1_funct_2(X56,X53,X54)|~r2_hidden(X56,k5_partfun1(X53,X54,X55))|(~v1_funct_1(X55)|~m1_subset_1(X55,k1_zfmisc_1(k2_zfmisc_1(X53,X54)))))&(m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(X53,X54)))|~r2_hidden(X56,k5_partfun1(X53,X54,X55))|(~v1_funct_1(X55)|~m1_subset_1(X55,k1_zfmisc_1(k2_zfmisc_1(X53,X54))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t158_funct_2])])])])).
cnf(c_0_95, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(esk1_0,k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0))))|~r1_tarski(k2_relat_1(esk3_0),esk3_0)), inference(spm,[status(thm)],[c_0_42, c_0_84])).
cnf(c_0_96, negated_conjecture, (k2_relat_1(esk3_0)=k2_zfmisc_1(esk1_0,esk2_0)|~v5_relat_1(esk3_0,esk3_0)|~v1_xboole_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_71])])).
cnf(c_0_97, plain, (v5_relat_1(X1,X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_57, c_0_87])).
cnf(c_0_98, negated_conjecture, (r1_tarski(esk3_0,X1)|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_88, c_0_60])).
fof(c_0_99, plain, ![X57, X58]:((~r1_tarski(X57,k2_zfmisc_1(X57,X58))|X57=k1_xboole_0)&(~r1_tarski(X57,k2_zfmisc_1(X58,X57))|X57=k1_xboole_0)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t135_zfmisc_1])])])).
cnf(c_0_100, negated_conjecture, (k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0))=X1|~v1_xboole_0(esk1_0)|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_89, c_0_58])).
cnf(c_0_101, negated_conjecture, (r1_tarski(k2_relat_1(esk3_0),X1)|~v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_49, c_0_90])).
cnf(c_0_102, negated_conjecture, (~r1_tarski(k5_partfun1(esk1_0,esk2_0,esk3_0),k1_funct_2(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_103, plain, (v1_xboole_0(X1)|r1_tarski(k5_partfun1(X1,X2,X3),X4)|~v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(spm,[status(thm)],[c_0_49, c_0_91])).
cnf(c_0_104, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_105, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_92])).
cnf(c_0_106, plain, (r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_107, plain, (X1=k1_xboole_0|r2_hidden(X2,k1_funct_2(X3,X1))|~v1_funct_1(X2)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_108, plain, (v1_funct_2(X1,X2,X3)|~r2_hidden(X1,k5_partfun1(X2,X3,X4))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_109, plain, (v1_funct_1(X1)|~r2_hidden(X1,k5_partfun1(X2,X3,X4))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_110, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r2_hidden(X1,k5_partfun1(X2,X3,X4))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_111, negated_conjecture, (k2_zfmisc_1(esk1_0,k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0)))=esk3_0|~r1_tarski(k2_zfmisc_1(esk1_0,k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0))),esk3_0)|~r1_tarski(k2_relat_1(esk3_0),esk3_0)), inference(spm,[status(thm)],[c_0_51, c_0_95])).
cnf(c_0_112, negated_conjecture, (k2_relat_1(esk3_0)=k2_zfmisc_1(esk1_0,esk2_0)|~v1_xboole_0(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_98])).
cnf(c_0_113, plain, (X1=k1_xboole_0|~r1_tarski(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_114, negated_conjecture, (k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0))=k2_relat_1(esk3_0)|~v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_115, negated_conjecture, (v1_xboole_0(esk1_0)|~v1_xboole_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_27]), c_0_104])])).
cnf(c_0_116, plain, (v1_xboole_0(X1)|~r1_tarski(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_105]), c_0_82])])).
cnf(c_0_117, plain, (X1=k1_xboole_0|r1_tarski(X2,k1_funct_2(X3,X1))|~v1_funct_2(esk5_2(X2,k1_funct_2(X3,X1)),X3,X1)|~m1_subset_1(esk5_2(X2,k1_funct_2(X3,X1)),k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v1_funct_1(esk5_2(X2,k1_funct_2(X3,X1)))), inference(spm,[status(thm)],[c_0_106, c_0_107])).
cnf(c_0_118, plain, (v1_funct_2(esk5_2(k5_partfun1(X1,X2,X3),X4),X1,X2)|r1_tarski(k5_partfun1(X1,X2,X3),X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(spm,[status(thm)],[c_0_108, c_0_38])).
cnf(c_0_119, plain, (r1_tarski(k5_partfun1(X1,X2,X3),X4)|v1_funct_1(esk5_2(k5_partfun1(X1,X2,X3),X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(spm,[status(thm)],[c_0_109, c_0_38])).
cnf(c_0_120, plain, (r1_tarski(k5_partfun1(X1,X2,X3),X4)|m1_subset_1(esk5_2(k5_partfun1(X1,X2,X3),X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(spm,[status(thm)],[c_0_110, c_0_38])).
cnf(c_0_121, negated_conjecture, (k2_zfmisc_1(esk1_0,k2_zfmisc_1(esk1_0,k2_zfmisc_1(esk1_0,esk2_0)))=esk3_0|~v1_xboole_0(esk2_0)|~r1_tarski(k2_zfmisc_1(esk1_0,k2_zfmisc_1(esk1_0,k2_zfmisc_1(esk1_0,esk2_0))),esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_77])).
cnf(c_0_122, negated_conjecture, (esk1_0=k1_xboole_0|~v1_xboole_0(esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_49])).
cnf(c_0_123, negated_conjecture, (v1_xboole_0(esk1_0)|~r1_tarski(esk2_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_115, c_0_116])).
cnf(c_0_124, plain, (X1=k1_xboole_0|r1_tarski(k5_partfun1(X2,X1,X3),k1_funct_2(X2,X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_119]), c_0_120])).
fof(c_0_125, plain, ![X59]:r1_tarski(k1_xboole_0,X59), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_126, negated_conjecture, (k2_zfmisc_1(esk1_0,k2_zfmisc_1(esk1_0,k2_zfmisc_1(esk1_0,esk2_0)))=esk3_0|~v1_xboole_0(k2_zfmisc_1(esk1_0,k2_zfmisc_1(esk1_0,esk2_0)))|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_121, c_0_77])).
fof(c_0_127, plain, ![X60, X61]:(~v1_funct_1(X61)|(~v1_funct_2(X61,X60,X60)|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(X60,X60))))|r2_hidden(X61,k1_funct_2(X60,X60))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_funct_2])])).
cnf(c_0_128, negated_conjecture, (esk1_0=k1_xboole_0|~r1_tarski(esk2_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_129, negated_conjecture, (esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_124]), c_0_27]), c_0_104])])).
cnf(c_0_130, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_125])).
cnf(c_0_131, negated_conjecture, (k2_zfmisc_1(esk1_0,k2_zfmisc_1(esk1_0,k2_zfmisc_1(esk1_0,esk2_0)))=esk3_0|~v1_xboole_0(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_66]), c_0_66])).
cnf(c_0_132, plain, (r2_hidden(X1,k1_funct_2(X2,X2))|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_133, negated_conjecture, (esk1_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_128, c_0_129]), c_0_130])])).
cnf(c_0_134, negated_conjecture, (esk3_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_131, c_0_129]), c_0_105]), c_0_105]), c_0_105]), c_0_129]), c_0_82])])).
cnf(c_0_135, plain, (r1_tarski(X1,k1_funct_2(X2,X2))|~v1_funct_2(esk5_2(X1,k1_funct_2(X2,X2)),X2,X2)|~m1_subset_1(esk5_2(X1,k1_funct_2(X2,X2)),k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~v1_funct_1(esk5_2(X1,k1_funct_2(X2,X2)))), inference(spm,[status(thm)],[c_0_106, c_0_132])).
fof(c_0_136, plain, ![X62]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X62)), inference(variable_rename,[status(thm)],[t4_subset_1])).
cnf(c_0_137, negated_conjecture, (~r1_tarski(k5_partfun1(k1_xboole_0,k1_xboole_0,k1_xboole_0),k1_funct_2(k1_xboole_0,k1_xboole_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_102, c_0_129]), c_0_129]), c_0_133]), c_0_134]), c_0_133])).
cnf(c_0_138, plain, (r1_tarski(k5_partfun1(X1,X1,X2),k1_funct_2(X1,X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_118]), c_0_119]), c_0_120])).
cnf(c_0_139, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_136])).
cnf(c_0_140, negated_conjecture, (v1_funct_1(k1_xboole_0)), inference(rw,[status(thm)],[c_0_104, c_0_134])).
cnf(c_0_141, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_138]), c_0_87]), c_0_139]), c_0_140])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 142
# Proof object clause steps            : 98
# Proof object formula steps           : 44
# Proof object conjectures             : 45
# Proof object clause conjectures      : 42
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 32
# Proof object initial formulas used   : 22
# Proof object generating inferences   : 59
# Proof object simplifying inferences  : 49
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 39
# Processed clauses                    : 5491
# ...of these trivial                  : 147
# ...subsumed                          : 4237
# ...remaining for further processing  : 1107
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 113
# Backward-rewritten                   : 678
# Generated clauses                    : 14708
# ...of the previous two non-trivial   : 12745
# Contextual simplify-reflections      : 65
# Paramodulations                      : 14703
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 273
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 254
# Current number of unprocessed clauses: 7139
# ...number of literals in the above   : 33131
# Current number of archived formulas  : 0
# Current number of archived clauses   : 829
# Clause-clause subsumption calls (NU) : 191422
# Rec. Clause-clause subsumption calls : 123026
# Non-unit clause-clause subsumptions  : 4390
# Unit Clause-clause subsumption calls : 1050
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 186873

# -------------------------------------------------
# User time                : 0.186 s
# System time              : 0.004 s
# Total time               : 0.189 s
# Maximum resident set size: 28688 pages
