# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/4132083f664a40b295170328a6323d8a.E_input', l32_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/4132083f664a40b295170328a6323d8a.E_input', t28_xboole_1)).
fof(d2_xboole_0, axiom, k1_xboole_0=o_0_0_xboole_0, file('data/MPTP/proofs/4132083f664a40b295170328a6323d8a.E_input', d2_xboole_0)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/4132083f664a40b295170328a6323d8a.E_input', t36_xboole_1)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/4132083f664a40b295170328a6323d8a.E_input', t41_xboole_1)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/4132083f664a40b295170328a6323d8a.E_input', t17_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/4132083f664a40b295170328a6323d8a.E_input', commutativity_k3_xboole_0)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/4132083f664a40b295170328a6323d8a.E_input', t39_xboole_1)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/4132083f664a40b295170328a6323d8a.E_input', t1_boole)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/4132083f664a40b295170328a6323d8a.E_input', d5_xboole_0)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/4132083f664a40b295170328a6323d8a.E_input', idempotence_k3_xboole_0)).
fof(t16_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3)), file('data/MPTP/proofs/4132083f664a40b295170328a6323d8a.E_input', t16_xboole_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/4132083f664a40b295170328a6323d8a.E_input', t3_xboole_0)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/4132083f664a40b295170328a6323d8a.E_input', commutativity_k2_xboole_0)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/4132083f664a40b295170328a6323d8a.E_input', d7_xboole_0)).
fof(t40_xboole_1, axiom, ![X1, X2]:k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/4132083f664a40b295170328a6323d8a.E_input', t40_xboole_1)).
fof(t47_xboole_1, conjecture, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/4132083f664a40b295170328a6323d8a.E_input', t47_xboole_1)).
fof(c_0_17, plain, ![X26, X27]:((k4_xboole_0(X26,X27)!=k1_xboole_0|r1_tarski(X26,X27))&(~r1_tarski(X26,X27)|k4_xboole_0(X26,X27)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
fof(c_0_18, plain, ![X47, X48]:(~r1_tarski(X47,X48)|k3_xboole_0(X47,X48)=X47), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_19, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20, plain, (k1_xboole_0=o_0_0_xboole_0), inference(split_conjunct,[status(thm)],[d2_xboole_0])).
cnf(c_0_21, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_22, plain, ![X16, X17]:r1_tarski(k4_xboole_0(X16,X17),X16), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_23, plain, ![X37, X38, X39]:k4_xboole_0(k4_xboole_0(X37,X38),X39)=k4_xboole_0(X37,k2_xboole_0(X38,X39)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
fof(c_0_24, plain, ![X42, X43]:r1_tarski(k3_xboole_0(X42,X43),X42), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_25, plain, ![X18, X19]:k3_xboole_0(X18,X19)=k3_xboole_0(X19,X18), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_26, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=o_0_0_xboole_0), inference(rw,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_28, plain, (k4_xboole_0(X1,X2)=o_0_0_xboole_0|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_21, c_0_20])).
cnf(c_0_29, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_31, plain, ![X40, X41]:k2_xboole_0(X40,k4_xboole_0(X41,X40))=k2_xboole_0(X40,X41), inference(variable_rename,[status(thm)],[t39_xboole_1])).
fof(c_0_32, plain, ![X9]:k2_xboole_0(X9,k1_xboole_0)=X9, inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_33, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, plain, (k3_xboole_0(X1,X2)=X1|k4_xboole_0(X1,X2)!=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_36, plain, (k4_xboole_0(X1,k2_xboole_0(X2,X1))=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])).
cnf(c_0_37, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_39, plain, ![X28, X29, X30, X31, X32, X33, X34, X35]:((((r2_hidden(X31,X28)|~r2_hidden(X31,X30)|X30!=k4_xboole_0(X28,X29))&(~r2_hidden(X31,X29)|~r2_hidden(X31,X30)|X30!=k4_xboole_0(X28,X29)))&(~r2_hidden(X32,X28)|r2_hidden(X32,X29)|r2_hidden(X32,X30)|X30!=k4_xboole_0(X28,X29)))&((~r2_hidden(esk4_3(X33,X34,X35),X35)|(~r2_hidden(esk4_3(X33,X34,X35),X33)|r2_hidden(esk4_3(X33,X34,X35),X34))|X35=k4_xboole_0(X33,X34))&((r2_hidden(esk4_3(X33,X34,X35),X33)|r2_hidden(esk4_3(X33,X34,X35),X35)|X35=k4_xboole_0(X33,X34))&(~r2_hidden(esk4_3(X33,X34,X35),X34)|r2_hidden(esk4_3(X33,X34,X35),X35)|X35=k4_xboole_0(X33,X34))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_40, plain, ![X46]:k3_xboole_0(X46,X46)=X46, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_41, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_42, plain, (k3_xboole_0(X1,k2_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_43, plain, (k2_xboole_0(X1,k4_xboole_0(X2,k2_xboole_0(X3,X1)))=k2_xboole_0(X1,k4_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_37, c_0_30])).
cnf(c_0_44, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X1)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_28, c_0_33])).
cnf(c_0_45, plain, (k2_xboole_0(X1,o_0_0_xboole_0)=X1), inference(rw,[status(thm)],[c_0_38, c_0_20])).
fof(c_0_46, plain, ![X13, X14, X15]:k3_xboole_0(k3_xboole_0(X13,X14),X15)=k3_xboole_0(X13,k3_xboole_0(X14,X15)), inference(variable_rename,[status(thm)],[t16_xboole_1])).
cnf(c_0_47, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_48, plain, ![X20, X21, X23, X24, X25]:(((r2_hidden(esk3_2(X20,X21),X20)|r1_xboole_0(X20,X21))&(r2_hidden(esk3_2(X20,X21),X21)|r1_xboole_0(X20,X21)))&(~r2_hidden(X25,X23)|~r2_hidden(X25,X24)|~r1_xboole_0(X23,X24))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_49, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_50, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_51, plain, (k2_xboole_0(X1,k4_xboole_0(k3_xboole_0(k2_xboole_0(X2,X1),X3),X2))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])).
fof(c_0_52, plain, ![X44, X45]:k2_xboole_0(X44,X45)=k2_xboole_0(X45,X44), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_53, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_54, plain, ![X7, X8]:((~r1_xboole_0(X7,X8)|k3_xboole_0(X7,X8)=k1_xboole_0)&(k3_xboole_0(X7,X8)!=k1_xboole_0|r1_xboole_0(X7,X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_55, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_47])).
cnf(c_0_56, plain, (r2_hidden(esk3_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_57, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_33, c_0_49])).
cnf(c_0_58, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(k2_xboole_0(X1,X2),X3),X1),X2)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_59, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_60, plain, (k3_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_53, c_0_49])).
cnf(c_0_61, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_62, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r2_hidden(esk3_2(X1,k4_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_63, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_64, plain, (k4_xboole_0(X1,X1)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_28, c_0_57])).
cnf(c_0_65, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(k2_xboole_0(X1,X2),X3),X2),X1)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_66, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_60, c_0_34])).
cnf(c_0_67, plain, (k4_xboole_0(X1,k2_xboole_0(X2,k2_xboole_0(X3,k4_xboole_0(X1,X2))))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_30, c_0_36])).
cnf(c_0_68, plain, (k3_xboole_0(X1,X2)=o_0_0_xboole_0|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_61, c_0_20])).
cnf(c_0_69, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_70, plain, (r1_tarski(o_0_0_xboole_0,X1)), inference(spm,[status(thm)],[c_0_29, c_0_64])).
cnf(c_0_71, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,k2_xboole_0(X2,X3)),X3),X2)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_72, plain, (k3_xboole_0(X1,k2_xboole_0(X2,k2_xboole_0(X3,k4_xboole_0(X1,X2))))=X1), inference(spm,[status(thm)],[c_0_35, c_0_67])).
cnf(c_0_73, plain, (k3_xboole_0(X1,k3_xboole_0(X2,k3_xboole_0(X1,X3)))=k3_xboole_0(X2,k3_xboole_0(X1,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_66]), c_0_53]), c_0_53])).
cnf(c_0_74, plain, (k3_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_29]), c_0_34])).
cnf(c_0_75, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X1))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_76, plain, (k3_xboole_0(o_0_0_xboole_0,X1)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_26, c_0_70])).
cnf(c_0_77, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(X2,k4_xboole_0(X1,X3))),X3)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_78, plain, (k2_xboole_0(o_0_0_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_45, c_0_59])).
cnf(c_0_79, plain, (r1_tarski(k3_xboole_0(X1,k3_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_41, c_0_53])).
cnf(c_0_80, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(X2,X3)),k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_81, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X2)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_28, c_0_41])).
cnf(c_0_82, plain, (k3_xboole_0(X1,k3_xboole_0(X2,k4_xboole_0(X1,X3)))=k3_xboole_0(X2,k4_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_83, plain, (k3_xboole_0(X1,k3_xboole_0(k4_xboole_0(X2,X1),X3))=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_75]), c_0_76])).
cnf(c_0_84, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_85, plain, (r1_tarski(k3_xboole_0(X1,X2),k2_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_79, c_0_42])).
cnf(c_0_86, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(X2,X3)),k4_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_80, c_0_59])).
cnf(c_0_87, plain, (k2_xboole_0(X1,k3_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_81]), c_0_45])).
cnf(c_0_88, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X3))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_89, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X1)))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_84]), c_0_34])).
cnf(c_0_90, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,k2_xboole_0(X2,X3)),X2),X3)), inference(spm,[status(thm)],[c_0_58, c_0_66])).
cnf(c_0_91, plain, (r1_tarski(X1,k2_xboole_0(X2,k2_xboole_0(X3,X1)))), inference(spm,[status(thm)],[c_0_85, c_0_42])).
cnf(c_0_92, plain, (r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X3,X2)))), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_93, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,X1)))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_94, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,k2_xboole_0(X2,X3)),X2),k4_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_90, c_0_37])).
cnf(c_0_95, plain, (k3_xboole_0(X1,k2_xboole_0(X2,k2_xboole_0(X3,X1)))=X1), inference(spm,[status(thm)],[c_0_26, c_0_91])).
cnf(c_0_96, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X3,X2)))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_26, c_0_92])).
cnf(c_0_97, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_93]), c_0_74])).
fof(c_0_98, plain, ![X10, X11]:k4_xboole_0(k2_xboole_0(X10,X11),X11)=k4_xboole_0(X10,X11), inference(variable_rename,[status(thm)],[t40_xboole_1])).
cnf(c_0_99, plain, (r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(k2_xboole_0(X3,X1),X2))), inference(spm,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_100, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_44]), c_0_45])).
cnf(c_0_101, plain, (k4_xboole_0(X1,k3_xboole_0(X2,k4_xboole_0(X3,X1)))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_74])).
cnf(c_0_102, plain, (k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_103, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,X2),X3),k4_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_104, plain, (k4_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X3,X2))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_101, c_0_96])).
cnf(c_0_105, plain, (k4_xboole_0(k2_xboole_0(X1,X2),k4_xboole_0(X2,X1))=k4_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_102, c_0_37])).
cnf(c_0_106, plain, (r1_tarski(k3_xboole_0(X1,X2),k4_xboole_0(X1,k4_xboole_0(X3,X2)))), inference(spm,[status(thm)],[c_0_103, c_0_104])).
fof(c_0_107, negated_conjecture, ~(![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(assume_negation,[status(cth)],[t47_xboole_1])).
cnf(c_0_108, plain, (k4_xboole_0(k2_xboole_0(X1,X2),k4_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[c_0_105, c_0_97])).
cnf(c_0_109, plain, (k2_xboole_0(X1,k2_xboole_0(X2,X1))=k2_xboole_0(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_36]), c_0_45]), c_0_59])).
cnf(c_0_110, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X3,X1)))=k3_xboole_0(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_106]), c_0_53]), c_0_82])).
fof(c_0_111, negated_conjecture, k4_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0))!=k4_xboole_0(esk1_0,esk2_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_107])])])).
cnf(c_0_112, plain, (k4_xboole_0(k2_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_102]), c_0_109])).
cnf(c_0_113, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_89, c_0_110])).
cnf(c_0_114, plain, (k2_xboole_0(X1,k4_xboole_0(X1,X2))=X1), inference(spm,[status(thm)],[c_0_100, c_0_74])).
cnf(c_0_115, negated_conjecture, (k4_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0))!=k4_xboole_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_116, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_114])).
cnf(c_0_117, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_115, c_0_116])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 118
# Proof object clause steps            : 84
# Proof object formula steps           : 34
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 57
# Proof object simplifying inferences  : 26
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 10111
# ...of these trivial                  : 4057
# ...subsumed                          : 4574
# ...remaining for further processing  : 1480
# Other redundant clauses eliminated   : 2050
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 14
# Backward-rewritten                   : 106
# Generated clauses                    : 477581
# ...of the previous two non-trivial   : 187824
# Contextual simplify-reflections      : 1
# Paramodulations                      : 475525
# Factorizations                       : 6
# Equation resolutions                 : 2050
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1350
#    Positive orientable unit clauses  : 1088
#    Positive unorientable unit clauses: 7
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 253
# Current number of unprocessed clauses: 177162
# ...number of literals in the above   : 262305
# Current number of archived formulas  : 0
# Current number of archived clauses   : 120
# Clause-clause subsumption calls (NU) : 24360
# Rec. Clause-clause subsumption calls : 18230
# Non-unit clause-clause subsumptions  : 2454
# Unit Clause-clause subsumption calls : 343
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8924
# BW rewrite match successes           : 443
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3772671

# -------------------------------------------------
# User time                : 2.302 s
# System time              : 0.070 s
# Total time               : 2.372 s
# Maximum resident set size: 155424 pages
