# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/413d7848000545f5b4be2de461164cf5.E_input', d1_tarski)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/413d7848000545f5b4be2de461164cf5.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/413d7848000545f5b4be2de461164cf5.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/413d7848000545f5b4be2de461164cf5.E_input', t71_enumset1)).
fof(t148_zfmisc_1, conjecture, ![X1, X2, X3, X4]:((r1_tarski(X1,X2)&(k3_xboole_0(X2,X3)=k1_tarski(X4)&r2_hidden(X4,X1)))=>k3_xboole_0(X1,X3)=k1_tarski(X4)), file('data/MPTP/proofs/413d7848000545f5b4be2de461164cf5.E_input', t148_zfmisc_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/413d7848000545f5b4be2de461164cf5.E_input', t48_xboole_1)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/413d7848000545f5b4be2de461164cf5.E_input', d4_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/413d7848000545f5b4be2de461164cf5.E_input', d3_tarski)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/413d7848000545f5b4be2de461164cf5.E_input', d2_tarski)).
fof(c_0_9, plain, ![X5, X6, X7, X8, X9, X10]:(((~r2_hidden(X7,X6)|X7=X5|X6!=k1_tarski(X5))&(X8!=X5|r2_hidden(X8,X6)|X6!=k1_tarski(X5)))&((~r2_hidden(esk1_2(X9,X10),X10)|esk1_2(X9,X10)!=X9|X10=k1_tarski(X9))&(r2_hidden(esk1_2(X9,X10),X10)|esk1_2(X9,X10)=X9|X10=k1_tarski(X9)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_10, plain, ![X12]:k2_tarski(X12,X12)=k1_tarski(X12), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_11, plain, ![X13, X14]:k1_enumset1(X13,X13,X14)=k2_tarski(X13,X14), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_12, plain, ![X15, X16, X17]:k2_enumset1(X15,X15,X16,X17)=k1_enumset1(X15,X16,X17), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_13, negated_conjecture, ~(![X1, X2, X3, X4]:((r1_tarski(X1,X2)&(k3_xboole_0(X2,X3)=k1_tarski(X4)&r2_hidden(X4,X1)))=>k3_xboole_0(X1,X3)=k1_tarski(X4))), inference(assume_negation,[status(cth)],[t148_zfmisc_1])).
cnf(c_0_14, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_18, negated_conjecture, ((r1_tarski(esk2_0,esk3_0)&(k3_xboole_0(esk3_0,esk4_0)=k1_tarski(esk5_0)&r2_hidden(esk5_0,esk2_0)))&k3_xboole_0(esk2_0,esk4_0)!=k1_tarski(esk5_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
fof(c_0_19, plain, ![X22, X23]:k4_xboole_0(X22,k4_xboole_0(X22,X23))=k3_xboole_0(X22,X23), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_20, plain, (X1=X3|X2!=k2_enumset1(X3,X3,X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_17])).
cnf(c_0_21, negated_conjecture, (k3_xboole_0(esk3_0,esk4_0)=k1_tarski(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_23, plain, ![X24, X25, X26, X27, X28, X29, X30, X31]:((((r2_hidden(X27,X24)|~r2_hidden(X27,X26)|X26!=k3_xboole_0(X24,X25))&(r2_hidden(X27,X25)|~r2_hidden(X27,X26)|X26!=k3_xboole_0(X24,X25)))&(~r2_hidden(X28,X24)|~r2_hidden(X28,X25)|r2_hidden(X28,X26)|X26!=k3_xboole_0(X24,X25)))&((~r2_hidden(esk6_3(X29,X30,X31),X31)|(~r2_hidden(esk6_3(X29,X30,X31),X29)|~r2_hidden(esk6_3(X29,X30,X31),X30))|X31=k3_xboole_0(X29,X30))&((r2_hidden(esk6_3(X29,X30,X31),X29)|r2_hidden(esk6_3(X29,X30,X31),X31)|X31=k3_xboole_0(X29,X30))&(r2_hidden(esk6_3(X29,X30,X31),X30)|r2_hidden(esk6_3(X29,X30,X31),X31)|X31=k3_xboole_0(X29,X30))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_24, plain, (X1=X2|~r2_hidden(X1,k2_enumset1(X2,X2,X2,X2))), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_25, negated_conjecture, (k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,esk4_0))=k2_enumset1(esk5_0,esk5_0,esk5_0,esk5_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_15]), c_0_16]), c_0_17]), c_0_22])).
fof(c_0_26, plain, ![X33, X34, X35, X36, X37]:((~r1_tarski(X33,X34)|(~r2_hidden(X35,X33)|r2_hidden(X35,X34)))&((r2_hidden(esk7_2(X36,X37),X36)|r1_tarski(X36,X37))&(~r2_hidden(esk7_2(X36,X37),X37)|r1_tarski(X36,X37)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_27, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28, negated_conjecture, (X1=esk5_0|~r2_hidden(X1,k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,esk4_0)))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29, plain, (r2_hidden(esk7_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_30, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_31, negated_conjecture, (r1_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_32, plain, (r2_hidden(X1,X2)|X3!=k4_xboole_0(X2,k4_xboole_0(X2,X4))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_27, c_0_22])).
cnf(c_0_33, plain, (r1_tarski(X1,X2)|~r2_hidden(esk7_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, negated_conjecture, (esk7_2(k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,esk4_0)),X1)=esk5_0|r1_tarski(k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,esk4_0)),X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X4!=k3_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_36, plain, ![X39, X40, X41, X42, X43, X44, X45, X46]:(((~r2_hidden(X42,X41)|(X42=X39|X42=X40)|X41!=k2_tarski(X39,X40))&((X43!=X39|r2_hidden(X43,X41)|X41!=k2_tarski(X39,X40))&(X43!=X40|r2_hidden(X43,X41)|X41!=k2_tarski(X39,X40))))&(((esk8_3(X44,X45,X46)!=X44|~r2_hidden(esk8_3(X44,X45,X46),X46)|X46=k2_tarski(X44,X45))&(esk8_3(X44,X45,X46)!=X45|~r2_hidden(esk8_3(X44,X45,X46),X46)|X46=k2_tarski(X44,X45)))&(r2_hidden(esk8_3(X44,X45,X46),X46)|(esk8_3(X44,X45,X46)=X44|esk8_3(X44,X45,X46)=X45)|X46=k2_tarski(X44,X45)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_37, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_38, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_39, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_40, negated_conjecture, (r1_tarski(k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,esk4_0)),X1)|~r2_hidden(esk5_0,X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_41, plain, (r2_hidden(X1,X4)|X4!=k4_xboole_0(X2,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_35, c_0_22])).
cnf(c_0_42, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_43, plain, (r2_hidden(X1,X2)|X3!=k4_xboole_0(X4,k4_xboole_0(X4,X2))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_37, c_0_22])).
cnf(c_0_44, negated_conjecture, (r1_tarski(X1,esk3_0)|~r2_hidden(esk7_2(X1,esk3_0),esk2_0)), inference(spm,[status(thm)],[c_0_33, c_0_38])).
cnf(c_0_45, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X3)|r2_hidden(esk7_2(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X3),X1)), inference(spm,[status(thm)],[c_0_39, c_0_29])).
cnf(c_0_46, negated_conjecture, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,esk4_0)))|~r2_hidden(esk5_0,X2)), inference(spm,[status(thm)],[c_0_30, c_0_40])).
cnf(c_0_47, plain, (r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_48, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_enumset1(X4,X4,X4,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_16]), c_0_17])).
cnf(c_0_49, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X3,k4_xboole_0(X3,X2)))), inference(er,[status(thm)],[c_0_43])).
cnf(c_0_50, negated_conjecture, (r1_tarski(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,X1)),esk3_0)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_51, negated_conjecture, (r2_hidden(X1,X2)|~r2_hidden(esk5_0,X2)|~r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_52, plain, (r2_hidden(X1,k2_enumset1(X2,X2,X2,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_48])])).
cnf(c_0_53, negated_conjecture, (X1=esk5_0|~r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_28, c_0_47])).
cnf(c_0_54, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X3)|r2_hidden(esk7_2(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X3),X2)), inference(spm,[status(thm)],[c_0_49, c_0_29])).
cnf(c_0_55, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,X2)))), inference(spm,[status(thm)],[c_0_30, c_0_50])).
cnf(c_0_56, plain, (r2_hidden(esk6_3(X1,X2,X3),X1)|r2_hidden(esk6_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_57, negated_conjecture, (r2_hidden(X1,k2_enumset1(X2,X2,X2,esk5_0))|~r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_58, negated_conjecture, (esk7_2(k4_xboole_0(X1,k4_xboole_0(X1,esk4_0)),X2)=esk5_0|r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,esk4_0)),X2)|~r2_hidden(esk7_2(k4_xboole_0(X1,k4_xboole_0(X1,esk4_0)),X2),esk3_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_59, negated_conjecture, (r1_tarski(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,X1)),X2)|r2_hidden(esk7_2(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,X1)),X2),esk3_0)), inference(spm,[status(thm)],[c_0_55, c_0_29])).
cnf(c_0_60, negated_conjecture, (r2_hidden(esk5_0,k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,esk4_0)))), inference(spm,[status(thm)],[c_0_52, c_0_25])).
cnf(c_0_61, plain, (X3=k4_xboole_0(X1,k4_xboole_0(X1,X2))|r2_hidden(esk6_3(X1,X2,X3),X3)|r2_hidden(esk6_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_56, c_0_22])).
cnf(c_0_62, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r2_hidden(esk6_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_63, negated_conjecture, (r1_tarski(X1,k2_enumset1(X2,X2,X2,esk5_0))|~r2_hidden(esk7_2(X1,k2_enumset1(X2,X2,X2,esk5_0)),esk4_0)|~r2_hidden(esk7_2(X1,k2_enumset1(X2,X2,X2,esk5_0)),esk3_0)), inference(spm,[status(thm)],[c_0_33, c_0_57])).
cnf(c_0_64, negated_conjecture, (esk7_2(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk4_0)),X1)=esk5_0|r1_tarski(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk4_0)),X1)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_65, negated_conjecture, (r2_hidden(esk5_0,esk4_0)), inference(spm,[status(thm)],[c_0_49, c_0_60])).
cnf(c_0_66, negated_conjecture, (r2_hidden(esk5_0,esk3_0)), inference(spm,[status(thm)],[c_0_39, c_0_60])).
cnf(c_0_67, negated_conjecture, (X1=k4_xboole_0(esk4_0,k4_xboole_0(esk4_0,X2))|esk6_3(esk4_0,X2,X1)=esk5_0|r2_hidden(esk6_3(esk4_0,X2,X1),X1)|~r2_hidden(esk6_3(esk4_0,X2,X1),esk3_0)), inference(spm,[status(thm)],[c_0_53, c_0_61])).
cnf(c_0_68, plain, (X3=k4_xboole_0(X1,k4_xboole_0(X1,X2))|r2_hidden(esk6_3(X1,X2,X3),X3)|r2_hidden(esk6_3(X1,X2,X3),X2)), inference(rw,[status(thm)],[c_0_62, c_0_22])).
cnf(c_0_69, negated_conjecture, (r1_tarski(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk4_0)),k2_enumset1(X1,X1,X1,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_65]), c_0_66])])).
cnf(c_0_70, plain, (X3=k3_xboole_0(X1,X2)|~r2_hidden(esk6_3(X1,X2,X3),X3)|~r2_hidden(esk6_3(X1,X2,X3),X1)|~r2_hidden(esk6_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_71, negated_conjecture, (X1=k4_xboole_0(esk4_0,k4_xboole_0(esk4_0,esk3_0))|esk6_3(esk4_0,esk3_0,X1)=esk5_0|r2_hidden(esk6_3(esk4_0,esk3_0,X1),X1)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_72, negated_conjecture, (r1_tarski(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk4_0)),k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,esk4_0)))), inference(spm,[status(thm)],[c_0_69, c_0_25])).
cnf(c_0_73, plain, (X3=k4_xboole_0(X1,k4_xboole_0(X1,X2))|~r2_hidden(esk6_3(X1,X2,X3),X3)|~r2_hidden(esk6_3(X1,X2,X3),X2)|~r2_hidden(esk6_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_70, c_0_22])).
cnf(c_0_74, negated_conjecture, (k4_xboole_0(esk4_0,k4_xboole_0(esk4_0,esk3_0))=k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,esk4_0))|esk6_3(esk4_0,esk3_0,k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,esk4_0)))=esk5_0), inference(spm,[status(thm)],[c_0_28, c_0_71])).
cnf(c_0_75, negated_conjecture, (k3_xboole_0(esk2_0,esk4_0)!=k1_tarski(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_76, negated_conjecture, (r2_hidden(X1,k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,esk4_0)))|~r2_hidden(X1,k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk4_0)))), inference(spm,[status(thm)],[c_0_30, c_0_72])).
cnf(c_0_77, negated_conjecture, (k4_xboole_0(esk4_0,k4_xboole_0(esk4_0,esk3_0))=k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_60]), c_0_66]), c_0_65])])).
cnf(c_0_78, negated_conjecture, (k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk4_0))!=k2_enumset1(esk5_0,esk5_0,esk5_0,esk5_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_15]), c_0_16]), c_0_17]), c_0_22])).
cnf(c_0_79, negated_conjecture, (X1=esk5_0|~r2_hidden(X1,k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk4_0)))), inference(spm,[status(thm)],[c_0_28, c_0_76])).
cnf(c_0_80, negated_conjecture, (X1=k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,esk4_0))|esk6_3(esk4_0,esk3_0,X1)=esk5_0|r2_hidden(esk6_3(esk4_0,esk3_0,X1),X1)), inference(rw,[status(thm)],[c_0_71, c_0_77])).
cnf(c_0_81, negated_conjecture, (k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,esk4_0))!=k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk4_0))), inference(rw,[status(thm)],[c_0_78, c_0_25])).
cnf(c_0_82, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X3,k4_xboole_0(X3,X4))|~r2_hidden(esk6_3(X3,X4,k4_xboole_0(X1,k4_xboole_0(X1,X2))),X4)|~r2_hidden(esk6_3(X3,X4,k4_xboole_0(X1,k4_xboole_0(X1,X2))),X3)|~r2_hidden(esk6_3(X3,X4,k4_xboole_0(X1,k4_xboole_0(X1,X2))),X2)|~r2_hidden(esk6_3(X3,X4,k4_xboole_0(X1,k4_xboole_0(X1,X2))),X1)), inference(spm,[status(thm)],[c_0_73, c_0_47])).
cnf(c_0_83, negated_conjecture, (esk6_3(esk4_0,esk3_0,k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk4_0)))=esk5_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_81])).
cnf(c_0_84, negated_conjecture, (r2_hidden(esk5_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_85, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_77]), c_0_66]), c_0_65]), c_0_84])]), c_0_81]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 86
# Proof object clause steps            : 67
# Proof object formula steps           : 19
# Proof object conjectures             : 39
# Proof object clause conjectures      : 36
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 31
# Proof object simplifying inferences  : 41
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 23
# Processed clauses                    : 1904
# ...of these trivial                  : 7
# ...subsumed                          : 989
# ...remaining for further processing  : 908
# Other redundant clauses eliminated   : 631
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 23
# Backward-rewritten                   : 26
# Generated clauses                    : 100469
# ...of the previous two non-trivial   : 97309
# Contextual simplify-reflections      : 10
# Paramodulations                      : 99666
# Factorizations                       : 174
# Equation resolutions                 : 632
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 829
#    Positive orientable unit clauses  : 52
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 776
# Current number of unprocessed clauses: 95423
# ...number of literals in the above   : 609015
# Current number of archived formulas  : 0
# Current number of archived clauses   : 75
# Clause-clause subsumption calls (NU) : 66289
# Rec. Clause-clause subsumption calls : 40761
# Non-unit clause-clause subsumptions  : 1022
# Unit Clause-clause subsumption calls : 1063
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 92
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2528105

# -------------------------------------------------
# User time                : 2.027 s
# System time              : 0.037 s
# Total time               : 2.064 s
# Maximum resident set size: 85020 pages
