# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t70_enumset1)).
fof(t138_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=>(k2_zfmisc_1(X1,X2)=k1_xboole_0|(r1_tarski(X1,X3)&r1_tarski(X2,X4)))), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t138_zfmisc_1)).
fof(t120_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&k2_zfmisc_1(X3,k2_xboole_0(X1,X2))=k2_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t120_zfmisc_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t75_enumset1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t4_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', commutativity_k3_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t28_xboole_1)).
fof(t123_zfmisc_1, axiom, ![X1, X2, X3, X4]:k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t123_zfmisc_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', commutativity_k2_tarski)).
fof(t125_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k4_xboole_0(X1,X2),X3)=k4_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&k2_zfmisc_1(X3,k4_xboole_0(X1,X2))=k4_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t125_zfmisc_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t100_xboole_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', idempotence_k3_xboole_0)).
fof(t22_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t22_xboole_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t7_xboole_0)).
fof(t77_xboole_1, axiom, ![X1, X2, X3]:~((~(r1_xboole_0(X1,X2))&(r1_tarski(X1,X3)&r1_xboole_0(X1,k3_xboole_0(X2,X3))))), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t77_xboole_1)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t2_boole)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t65_xboole_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t113_zfmisc_1)).
fof(t127_zfmisc_1, axiom, ![X1, X2, X3, X4]:((r1_xboole_0(X1,X2)|r1_xboole_0(X3,X4))=>r1_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t127_zfmisc_1)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', symmetry_r1_xboole_0)).
fof(t117_zfmisc_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&((r1_tarski(k2_zfmisc_1(X2,X1),k2_zfmisc_1(X3,X1))|r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))&~(r1_tarski(X2,X3))))), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t117_zfmisc_1)).
fof(t92_xboole_1, axiom, ![X1]:k5_xboole_0(X1,X1)=k1_xboole_0, file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t92_xboole_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', l32_xboole_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/416c50e8104b446dbb26a291a3a7d6bd.E_input', t2_xboole_1)).
fof(c_0_29, plain, ![X71, X72]:k3_tarski(k2_tarski(X71,X72))=k2_xboole_0(X71,X72), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_30, plain, ![X48, X49]:k1_enumset1(X48,X48,X49)=k2_tarski(X48,X49), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_31, negated_conjecture, ~(![X1, X2, X3, X4]:(r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=>(k2_zfmisc_1(X1,X2)=k1_xboole_0|(r1_tarski(X1,X3)&r1_tarski(X2,X4))))), inference(assume_negation,[status(cth)],[t138_zfmisc_1])).
fof(c_0_32, plain, ![X80, X81, X82]:(k2_zfmisc_1(k2_xboole_0(X80,X81),X82)=k2_xboole_0(k2_zfmisc_1(X80,X82),k2_zfmisc_1(X81,X82))&k2_zfmisc_1(X82,k2_xboole_0(X80,X81))=k2_xboole_0(k2_zfmisc_1(X82,X80),k2_zfmisc_1(X82,X81))), inference(variable_rename,[status(thm)],[t120_zfmisc_1])).
cnf(c_0_33, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_35, plain, ![X74, X75, X76]:k2_enumset1(X74,X74,X75,X76)=k1_enumset1(X74,X75,X76), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_36, plain, ![X24, X25, X26, X27]:k3_enumset1(X24,X24,X25,X26,X27)=k2_enumset1(X24,X25,X26,X27), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_37, plain, ![X29, X30, X31, X32, X33]:k4_enumset1(X29,X29,X30,X31,X32,X33)=k3_enumset1(X29,X30,X31,X32,X33), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_38, plain, ![X40, X41, X42, X43, X44, X45]:k5_enumset1(X40,X40,X41,X42,X43,X44,X45)=k4_enumset1(X40,X41,X42,X43,X44,X45), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_39, plain, ![X83, X84, X85, X86, X87, X88, X89]:k6_enumset1(X83,X83,X84,X85,X86,X87,X88,X89)=k5_enumset1(X83,X84,X85,X86,X87,X88,X89), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_40, plain, ![X63, X64, X66, X67, X68]:((r1_xboole_0(X63,X64)|r2_hidden(esk6_2(X63,X64),k3_xboole_0(X63,X64)))&(~r2_hidden(X68,k3_xboole_0(X66,X67))|~r1_xboole_0(X66,X67))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_41, plain, ![X22, X23]:k3_xboole_0(X22,X23)=k3_xboole_0(X23,X22), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_42, plain, ![X69, X70]:(~r1_tarski(X69,X70)|k3_xboole_0(X69,X70)=X69), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_43, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))&(k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0&(~r1_tarski(esk1_0,esk3_0)|~r1_tarski(esk2_0,esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])])).
cnf(c_0_44, plain, (k2_zfmisc_1(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_45, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_46, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_47, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_48, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_49, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_51, plain, ![X34, X35, X36, X37]:k2_zfmisc_1(k3_xboole_0(X34,X35),k3_xboole_0(X36,X37))=k3_xboole_0(k2_zfmisc_1(X34,X36),k2_zfmisc_1(X35,X37)), inference(variable_rename,[status(thm)],[t123_zfmisc_1])).
fof(c_0_52, plain, ![X77, X78]:k2_tarski(X77,X78)=k2_tarski(X78,X77), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_53, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_54, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_55, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_56, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_57, plain, ![X56, X57, X58]:(k2_zfmisc_1(k4_xboole_0(X56,X57),X58)=k4_xboole_0(k2_zfmisc_1(X56,X58),k2_zfmisc_1(X57,X58))&k2_zfmisc_1(X58,k4_xboole_0(X56,X57))=k4_xboole_0(k2_zfmisc_1(X58,X56),k2_zfmisc_1(X58,X57))), inference(variable_rename,[status(thm)],[t125_zfmisc_1])).
fof(c_0_58, plain, ![X19, X20]:k4_xboole_0(X19,X20)=k5_xboole_0(X19,k3_xboole_0(X19,X20)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_59, plain, ![X28]:k3_xboole_0(X28,X28)=X28, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_60, plain, (k2_zfmisc_1(X1,k3_tarski(k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)))=k3_tarski(k6_enumset1(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_50]), c_0_50])).
cnf(c_0_61, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_62, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
fof(c_0_63, plain, ![X38, X39]:k2_xboole_0(X38,k3_xboole_0(X38,X39))=X38, inference(variable_rename,[status(thm)],[t22_xboole_1])).
cnf(c_0_64, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_65, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))=k2_zfmisc_1(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
fof(c_0_66, plain, ![X17]:(X17=k1_xboole_0|r2_hidden(esk5_1(X17),X17)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
fof(c_0_67, plain, ![X53, X54, X55]:(r1_xboole_0(X53,X54)|(~r1_tarski(X53,X55)|~r1_xboole_0(X53,k3_xboole_0(X54,X55)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t77_xboole_1])])])).
fof(c_0_68, plain, ![X16]:k3_xboole_0(X16,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_69, plain, ![X79]:r1_xboole_0(X79,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
fof(c_0_70, plain, ![X46, X47]:((k2_zfmisc_1(X46,X47)!=k1_xboole_0|(X46=k1_xboole_0|X47=k1_xboole_0))&((X46!=k1_xboole_0|k2_zfmisc_1(X46,X47)=k1_xboole_0)&(X47!=k1_xboole_0|k2_zfmisc_1(X46,X47)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_71, plain, (k2_zfmisc_1(k4_xboole_0(X1,X2),X3)=k4_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_72, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_73, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_74, plain, (k2_zfmisc_1(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_75, plain, (k3_tarski(k6_enumset1(k2_zfmisc_1(k3_xboole_0(X1,X2),X3),k2_zfmisc_1(k3_xboole_0(X1,X2),X3),k2_zfmisc_1(k3_xboole_0(X1,X2),X3),k2_zfmisc_1(k3_xboole_0(X1,X2),X3),k2_zfmisc_1(k3_xboole_0(X1,X2),X3),k2_zfmisc_1(k3_xboole_0(X1,X2),X3),k2_zfmisc_1(k3_xboole_0(X1,X2),X3),k3_xboole_0(k2_zfmisc_1(X1,X4),k2_zfmisc_1(X2,X5))))=k2_zfmisc_1(k3_xboole_0(X1,X2),k3_tarski(k6_enumset1(X3,X3,X3,X3,X3,X3,X3,k3_xboole_0(X4,X5))))), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_76, plain, (k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2)=k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_34]), c_0_34]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_50]), c_0_50])).
cnf(c_0_77, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_78, negated_conjecture, (~r2_hidden(X1,k2_zfmisc_1(esk1_0,esk2_0))|~r1_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_79, plain, (X1=k1_xboole_0|r2_hidden(esk5_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_80, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_81, plain, ![X12, X13, X14, X15]:((~r1_xboole_0(X12,X13)|r1_xboole_0(k2_zfmisc_1(X12,X14),k2_zfmisc_1(X13,X15)))&(~r1_xboole_0(X14,X15)|r1_xboole_0(k2_zfmisc_1(X12,X14),k2_zfmisc_1(X13,X15)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t127_zfmisc_1])])])).
fof(c_0_82, plain, ![X61, X62]:(~r1_xboole_0(X61,X62)|r1_xboole_0(X62,X61)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
cnf(c_0_83, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,X3)|~r1_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_84, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_85, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_69])).
fof(c_0_86, plain, ![X50, X51, X52]:((~r1_tarski(k2_zfmisc_1(X51,X50),k2_zfmisc_1(X52,X50))|r1_tarski(X51,X52)|X50=k1_xboole_0)&(~r1_tarski(k2_zfmisc_1(X50,X51),k2_zfmisc_1(X50,X52))|r1_tarski(X51,X52)|X50=k1_xboole_0)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t117_zfmisc_1])])])])).
cnf(c_0_87, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_88, plain, (k2_zfmisc_1(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X3)=k5_xboole_0(k2_zfmisc_1(X1,X3),k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_72]), c_0_72])).
cnf(c_0_89, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))=k2_zfmisc_1(k3_xboole_0(X1,X3),X2)), inference(spm,[status(thm)],[c_0_61, c_0_73])).
cnf(c_0_90, plain, (k2_zfmisc_1(k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2)),X3)=k3_tarski(k6_enumset1(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X1,X3),k2_zfmisc_1(X1,X3),k2_zfmisc_1(X1,X3),k2_zfmisc_1(X1,X3),k2_zfmisc_1(X1,X3),k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_50]), c_0_50])).
cnf(c_0_91, negated_conjecture, (k3_tarski(k6_enumset1(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(k3_xboole_0(esk1_0,esk3_0),X1)))=k2_zfmisc_1(k3_xboole_0(esk1_0,esk3_0),k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,k3_xboole_0(esk2_0,esk4_0))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_65]), c_0_76])).
cnf(c_0_92, plain, (k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,k3_xboole_0(X1,X2)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_45]), c_0_46]), c_0_47]), c_0_48]), c_0_49]), c_0_50])).
fof(c_0_93, plain, ![X73]:k5_xboole_0(X73,X73)=k1_xboole_0, inference(variable_rename,[status(thm)],[t92_xboole_1])).
cnf(c_0_94, negated_conjecture, (~r1_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80])).
cnf(c_0_95, plain, (r1_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X4,X2))|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_96, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_97, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_85])])).
fof(c_0_98, plain, ![X59, X60]:((k4_xboole_0(X59,X60)!=k1_xboole_0|r1_tarski(X59,X60))&(~r1_tarski(X59,X60)|k4_xboole_0(X59,X60)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_99, plain, (r1_tarski(X2,X3)|X1=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_100, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_87])).
cnf(c_0_101, plain, (k5_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(k3_xboole_0(X1,X3),X2))=k2_zfmisc_1(k5_xboole_0(X1,k3_xboole_0(X1,X3)),X2)), inference(rw,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_102, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk1_0,esk3_0),esk2_0)=k2_zfmisc_1(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_92]), c_0_92])).
cnf(c_0_103, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_93])).
fof(c_0_104, plain, ![X21]:r1_tarski(k1_xboole_0,X21), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_105, negated_conjecture, (~r1_xboole_0(esk4_0,esk2_0)), inference(spm,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_106, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_107, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_108, plain, (X1=k1_xboole_0|r1_tarski(X2,k1_xboole_0)|~r1_tarski(k2_zfmisc_1(X1,X2),k1_xboole_0)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_109, negated_conjecture, (k2_zfmisc_1(k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk3_0)),esk2_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_103])).
cnf(c_0_110, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_111, negated_conjecture, (~r1_tarski(esk2_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_112, plain, (r1_tarski(X1,X2)|k5_xboole_0(X1,k3_xboole_0(X1,X2))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_107, c_0_72])).
cnf(c_0_113, negated_conjecture, (k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk3_0))=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_110])]), c_0_111])).
cnf(c_0_114, negated_conjecture, (r1_tarski(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_112, c_0_113])).
cnf(c_0_115, plain, (k2_zfmisc_1(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_116, negated_conjecture, (k3_xboole_0(esk1_0,esk3_0)=esk1_0), inference(spm,[status(thm)],[c_0_55, c_0_114])).
cnf(c_0_117, plain, (k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,k1_xboole_0))=X1), inference(spm,[status(thm)],[c_0_92, c_0_84])).
cnf(c_0_118, plain, (k2_zfmisc_1(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))=k5_xboole_0(k2_zfmisc_1(X1,X2),k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_115, c_0_72]), c_0_72])).
cnf(c_0_119, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))=k2_zfmisc_1(X1,k3_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_61, c_0_73])).
cnf(c_0_120, negated_conjecture, (k2_zfmisc_1(esk1_0,k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,k3_xboole_0(esk2_0,esk4_0))))=k2_zfmisc_1(esk1_0,k3_tarski(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_116]), c_0_60]), c_0_116])).
cnf(c_0_121, plain, (k3_tarski(k6_enumset1(k1_xboole_0,k1_xboole_0,k1_xboole_0,k1_xboole_0,k1_xboole_0,k1_xboole_0,k1_xboole_0,X1))=X1), inference(spm,[status(thm)],[c_0_117, c_0_76])).
cnf(c_0_122, negated_conjecture, (~r2_hidden(X1,k2_zfmisc_1(esk1_0,esk2_0))|~r1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_53, c_0_65])).
cnf(c_0_123, plain, (k5_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,k3_xboole_0(X2,X3)))=k2_zfmisc_1(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))), inference(rw,[status(thm)],[c_0_118, c_0_119])).
cnf(c_0_124, negated_conjecture, (k2_zfmisc_1(esk1_0,k3_xboole_0(esk2_0,esk4_0))=k2_zfmisc_1(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121]), c_0_117])).
cnf(c_0_125, negated_conjecture, (~r1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_79]), c_0_80])).
cnf(c_0_126, plain, (r1_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_127, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))|k2_zfmisc_1(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_119]), c_0_123])).
cnf(c_0_128, negated_conjecture, (k2_zfmisc_1(esk1_0,k5_xboole_0(esk2_0,k3_xboole_0(esk2_0,esk4_0)))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124]), c_0_103])).
cnf(c_0_129, negated_conjecture, (~r1_tarski(esk1_0,esk3_0)|~r1_tarski(esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_130, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk6_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_131, negated_conjecture, (~r1_xboole_0(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_125, c_0_126])).
cnf(c_0_132, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk1_0,esk4_0))), inference(spm,[status(thm)],[c_0_127, c_0_128])).
cnf(c_0_133, negated_conjecture, (~r1_tarski(esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_129, c_0_114])])).
cnf(c_0_134, negated_conjecture, (r2_hidden(esk6_2(esk1_0,esk3_0),esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_116]), c_0_131])).
cnf(c_0_135, negated_conjecture, (esk1_0=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_132]), c_0_133])).
cnf(c_0_136, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_84]), c_0_85])])).
cnf(c_0_137, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_134, c_0_135]), c_0_135]), c_0_136]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 138
# Proof object clause steps            : 79
# Proof object formula steps           : 59
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 35
# Proof object initial formulas used   : 29
# Proof object generating inferences   : 30
# Proof object simplifying inferences  : 77
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 29
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 8
# Initial clauses in saturation        : 31
# Processed clauses                    : 6130
# ...of these trivial                  : 75
# ...subsumed                          : 5123
# ...remaining for further processing  : 932
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 118
# Backward-rewritten                   : 331
# Generated clauses                    : 76239
# ...of the previous two non-trivial   : 64769
# Contextual simplify-reflections      : 12
# Paramodulations                      : 76237
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 450
#    Positive orientable unit clauses  : 70
#    Positive unorientable unit clauses: 11
#    Negative unit clauses             : 13
#    Non-unit-clauses                  : 356
# Current number of unprocessed clauses: 58281
# ...number of literals in the above   : 133204
# Current number of archived formulas  : 0
# Current number of archived clauses   : 488
# Clause-clause subsumption calls (NU) : 53688
# Rec. Clause-clause subsumption calls : 52755
# Non-unit clause-clause subsumptions  : 2241
# Unit Clause-clause subsumption calls : 1074
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 417
# BW rewrite match successes           : 160
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1530678

# -------------------------------------------------
# User time                : 0.655 s
# System time              : 0.013 s
# Total time               : 0.668 s
# Maximum resident set size: 49620 pages
