# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/41a2f19e63504149b236c471f556e9e1.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/41a2f19e63504149b236c471f556e9e1.E_input', t70_enumset1)).
fof(t43_zfmisc_1, conjecture, ![X1, X2, X3]:~((k1_tarski(X1)=k2_xboole_0(X2,X3)&(~((X2=k1_tarski(X1)&X3=k1_tarski(X1)))&(~((X2=k1_xboole_0&X3=k1_tarski(X1)))&~((X2=k1_tarski(X1)&X3=k1_xboole_0)))))), file('data/MPTP/proofs/41a2f19e63504149b236c471f556e9e1.E_input', t43_zfmisc_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/41a2f19e63504149b236c471f556e9e1.E_input', t7_xboole_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/41a2f19e63504149b236c471f556e9e1.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/41a2f19e63504149b236c471f556e9e1.E_input', t72_enumset1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/41a2f19e63504149b236c471f556e9e1.E_input', t69_enumset1)).
fof(d3_xboole_0, axiom, ![X1, X2, X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2)))), file('data/MPTP/proofs/41a2f19e63504149b236c471f556e9e1.E_input', d3_xboole_0)).
fof(l3_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(X1,k1_tarski(X2))<=>(X1=k1_xboole_0|X1=k1_tarski(X2))), file('data/MPTP/proofs/41a2f19e63504149b236c471f556e9e1.E_input', l3_zfmisc_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/41a2f19e63504149b236c471f556e9e1.E_input', d1_tarski)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/41a2f19e63504149b236c471f556e9e1.E_input', t7_xboole_0)).
fof(l22_zfmisc_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>k2_xboole_0(k1_tarski(X1),X2)=X2), file('data/MPTP/proofs/41a2f19e63504149b236c471f556e9e1.E_input', l22_zfmisc_1)).
fof(t11_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k2_xboole_0(X1,X2),X3)=>r1_tarski(X1,X3)), file('data/MPTP/proofs/41a2f19e63504149b236c471f556e9e1.E_input', t11_xboole_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/41a2f19e63504149b236c471f556e9e1.E_input', d1_xboole_0)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/41a2f19e63504149b236c471f556e9e1.E_input', t12_xboole_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/41a2f19e63504149b236c471f556e9e1.E_input', fc1_xboole_0)).
fof(c_0_16, plain, ![X22, X23]:k3_tarski(k2_tarski(X22,X23))=k2_xboole_0(X22,X23), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_17, plain, ![X19, X20]:k1_enumset1(X19,X19,X20)=k2_tarski(X19,X20), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_18, negated_conjecture, ~(![X1, X2, X3]:~((k1_tarski(X1)=k2_xboole_0(X2,X3)&(~((X2=k1_tarski(X1)&X3=k1_tarski(X1)))&(~((X2=k1_xboole_0&X3=k1_tarski(X1)))&~((X2=k1_tarski(X1)&X3=k1_xboole_0))))))), inference(assume_negation,[status(cth)],[t43_zfmisc_1])).
fof(c_0_19, plain, ![X17, X18]:r1_tarski(X17,k2_xboole_0(X17,X18)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_20, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_22, plain, ![X14, X15, X16]:k2_enumset1(X14,X14,X15,X16)=k1_enumset1(X14,X15,X16), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_23, plain, ![X10, X11, X12, X13]:k3_enumset1(X10,X10,X11,X12,X13)=k2_enumset1(X10,X11,X12,X13), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_24, negated_conjecture, (k1_tarski(esk1_0)=k2_xboole_0(esk2_0,esk3_0)&((esk2_0!=k1_tarski(esk1_0)|esk3_0!=k1_tarski(esk1_0))&((esk2_0!=k1_xboole_0|esk3_0!=k1_tarski(esk1_0))&(esk2_0!=k1_tarski(esk1_0)|esk3_0!=k1_xboole_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).
fof(c_0_25, plain, ![X21]:k2_tarski(X21,X21)=k1_tarski(X21), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_26, plain, ![X39, X40, X41, X42, X43, X44, X45, X46]:(((~r2_hidden(X42,X41)|(r2_hidden(X42,X39)|r2_hidden(X42,X40))|X41!=k2_xboole_0(X39,X40))&((~r2_hidden(X43,X39)|r2_hidden(X43,X41)|X41!=k2_xboole_0(X39,X40))&(~r2_hidden(X43,X40)|r2_hidden(X43,X41)|X41!=k2_xboole_0(X39,X40))))&(((~r2_hidden(esk7_3(X44,X45,X46),X44)|~r2_hidden(esk7_3(X44,X45,X46),X46)|X46=k2_xboole_0(X44,X45))&(~r2_hidden(esk7_3(X44,X45,X46),X45)|~r2_hidden(esk7_3(X44,X45,X46),X46)|X46=k2_xboole_0(X44,X45)))&(r2_hidden(esk7_3(X44,X45,X46),X46)|(r2_hidden(esk7_3(X44,X45,X46),X44)|r2_hidden(esk7_3(X44,X45,X46),X45))|X46=k2_xboole_0(X44,X45)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])).
fof(c_0_27, plain, ![X48, X49]:((~r1_tarski(X48,k1_tarski(X49))|(X48=k1_xboole_0|X48=k1_tarski(X49)))&((X48!=k1_xboole_0|r1_tarski(X48,k1_tarski(X49)))&(X48!=k1_tarski(X49)|r1_tarski(X48,k1_tarski(X49))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_zfmisc_1])])])).
cnf(c_0_28, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_30, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, negated_conjecture, (k1_tarski(esk1_0)=k2_xboole_0(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_35, plain, ![X28, X29, X30, X31, X32, X33]:(((~r2_hidden(X30,X29)|X30=X28|X29!=k1_tarski(X28))&(X31!=X28|r2_hidden(X31,X29)|X29!=k1_tarski(X28)))&((~r2_hidden(esk5_2(X32,X33),X33)|esk5_2(X32,X33)!=X32|X33=k1_tarski(X32))&(r2_hidden(esk5_2(X32,X33),X33)|esk5_2(X32,X33)=X32|X33=k1_tarski(X32)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_36, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|~r1_tarski(X1,k1_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (r1_tarski(X1,k3_tarski(k3_enumset1(X1,X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_29]), c_0_30]), c_0_31])).
cnf(c_0_38, negated_conjecture, (k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=k3_tarski(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_33]), c_0_21]), c_0_29]), c_0_30]), c_0_30]), c_0_31]), c_0_31])).
cnf(c_0_39, plain, (r2_hidden(X1,X3)|X3!=k3_tarski(k3_enumset1(X4,X4,X4,X4,X2))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_29]), c_0_30]), c_0_31])).
cnf(c_0_40, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41, plain, (X1=k1_xboole_0|X1=k3_enumset1(X2,X2,X2,X2,X2)|~r1_tarski(X1,k3_enumset1(X2,X2,X2,X2,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_33]), c_0_33]), c_0_21]), c_0_21]), c_0_30]), c_0_30]), c_0_31]), c_0_31])).
cnf(c_0_42, negated_conjecture, (r1_tarski(esk2_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_43, negated_conjecture, (r2_hidden(X1,X2)|X2!=k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_39, c_0_38])).
cnf(c_0_44, negated_conjecture, (esk2_0!=k1_tarski(esk1_0)|esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_45, plain, (X1=X3|X2!=k3_enumset1(X3,X3,X3,X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_33]), c_0_21]), c_0_30]), c_0_31])).
cnf(c_0_46, negated_conjecture, (k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=esk2_0|k1_xboole_0=esk2_0), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_47, negated_conjecture, (r2_hidden(X1,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))|~r2_hidden(X1,esk3_0)), inference(er,[status(thm)],[c_0_43])).
fof(c_0_48, plain, ![X35]:(X35=k1_xboole_0|r2_hidden(esk6_1(X35),X35)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_49, negated_conjecture, (esk3_0!=k1_xboole_0|esk2_0!=k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_33]), c_0_21]), c_0_30]), c_0_31])).
fof(c_0_50, plain, ![X8, X9]:(~r2_hidden(X8,X9)|k2_xboole_0(k1_tarski(X8),X9)=X9), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l22_zfmisc_1])])).
cnf(c_0_51, negated_conjecture, (k1_xboole_0=esk2_0|X1=esk1_0|X2!=esk2_0|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_52, negated_conjecture, (k1_xboole_0=esk2_0|r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_47, c_0_46])).
cnf(c_0_53, plain, (X1=k1_xboole_0|r2_hidden(esk6_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_54, negated_conjecture, (k1_xboole_0=esk2_0|esk3_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_49, c_0_46])).
cnf(c_0_55, plain, (k2_xboole_0(k1_tarski(X1),X2)=X2|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_56, negated_conjecture, (k1_xboole_0=esk2_0|X1=esk1_0|~r2_hidden(X1,esk2_0)), inference(er,[status(thm)],[c_0_51])).
cnf(c_0_57, negated_conjecture, (k1_xboole_0=esk2_0|r2_hidden(esk6_1(esk3_0),esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54])).
fof(c_0_58, plain, ![X50, X51, X52]:(~r1_tarski(k2_xboole_0(X50,X51),X52)|r1_tarski(X50,X52)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_xboole_1])])).
cnf(c_0_59, plain, (k3_tarski(k3_enumset1(k3_enumset1(X1,X1,X1,X1,X1),k3_enumset1(X1,X1,X1,X1,X1),k3_enumset1(X1,X1,X1,X1,X1),k3_enumset1(X1,X1,X1,X1,X1),X2))=X2|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_33]), c_0_21]), c_0_29]), c_0_30]), c_0_30]), c_0_30]), c_0_31]), c_0_31]), c_0_31]), c_0_31])).
cnf(c_0_60, negated_conjecture, (esk6_1(esk3_0)=esk1_0|k1_xboole_0=esk2_0), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_61, negated_conjecture, (esk2_0!=k1_tarski(esk1_0)|esk3_0!=k1_tarski(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_62, plain, (r1_tarski(X1,X3)|~r1_tarski(k2_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_63, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_64, plain, ![X24, X25, X26]:((~v1_xboole_0(X24)|~r2_hidden(X25,X24))&(r2_hidden(esk4_1(X26),X26)|v1_xboole_0(X26))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_65, plain, (r2_hidden(esk5_2(X1,X2),X2)|esk5_2(X1,X2)=X1|X2=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_66, negated_conjecture, (k3_tarski(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,X1))=X1|k1_xboole_0=esk2_0|~r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_59, c_0_46])).
cnf(c_0_67, negated_conjecture, (k1_xboole_0=esk2_0|r2_hidden(esk1_0,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_60]), c_0_54])).
cnf(c_0_68, negated_conjecture, (esk2_0!=k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)|esk3_0!=k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_33]), c_0_33]), c_0_21]), c_0_21]), c_0_30]), c_0_30]), c_0_31]), c_0_31])).
fof(c_0_69, plain, ![X37, X38]:(~r1_tarski(X37,X38)|k2_xboole_0(X37,X38)=X38), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_70, plain, (r1_tarski(X1,X3)|~r1_tarski(k3_tarski(k3_enumset1(X1,X1,X1,X1,X2)),X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_29]), c_0_30]), c_0_31])).
cnf(c_0_71, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k3_enumset1(X2,X2,X2,X2,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_33]), c_0_21]), c_0_30]), c_0_31])).
cnf(c_0_72, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_73, plain, (esk5_2(X1,X2)=X1|X2=k3_enumset1(X1,X1,X1,X1,X1)|r2_hidden(esk5_2(X1,X2),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_33]), c_0_21]), c_0_30]), c_0_31])).
cnf(c_0_74, negated_conjecture, (k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=esk3_0|k1_xboole_0=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_38])).
cnf(c_0_75, negated_conjecture, (k1_xboole_0=esk2_0|esk3_0!=esk2_0), inference(spm,[status(thm)],[c_0_68, c_0_46])).
cnf(c_0_76, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_77, negated_conjecture, (r1_tarski(esk2_0,X1)|~r1_tarski(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),X1)), inference(spm,[status(thm)],[c_0_70, c_0_38])).
cnf(c_0_78, negated_conjecture, (esk2_0!=k1_xboole_0|esk3_0!=k1_tarski(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_79, plain, (r2_hidden(X1,X2)|X2!=k3_enumset1(X1,X1,X1,X1,X1)), inference(er,[status(thm)],[c_0_71])).
cnf(c_0_80, plain, (X1=k3_enumset1(X2,X2,X2,X2,X2)|esk5_2(X2,X1)=X2|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_81, plain, (X1=X2|~r2_hidden(X1,k3_enumset1(X2,X2,X2,X2,X2))), inference(er,[status(thm)],[c_0_45])).
cnf(c_0_82, negated_conjecture, (k1_xboole_0=esk2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_74]), c_0_75])).
cnf(c_0_83, plain, (k3_tarski(k3_enumset1(X1,X1,X1,X1,X2))=X2|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_29]), c_0_30]), c_0_31])).
cnf(c_0_84, negated_conjecture, (r1_tarski(esk2_0,k3_tarski(k3_enumset1(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),X1)))), inference(spm,[status(thm)],[c_0_77, c_0_37])).
cnf(c_0_85, negated_conjecture, (esk2_0!=k1_xboole_0|esk3_0!=k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_33]), c_0_21]), c_0_30]), c_0_31])).
cnf(c_0_86, plain, (X2=k1_tarski(X1)|~r2_hidden(esk5_2(X1,X2),X2)|esk5_2(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_87, plain, (esk5_2(X1,X2)=X1|~v1_xboole_0(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80])]), c_0_72])).
cnf(c_0_88, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_89, negated_conjecture, (X1=esk1_0|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_81, c_0_47])).
cnf(c_0_90, plain, (X1=esk2_0|r2_hidden(esk6_1(X1),X1)), inference(rw,[status(thm)],[c_0_53, c_0_82])).
cnf(c_0_91, negated_conjecture, (k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=esk3_0|~r1_tarski(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_38, c_0_83])).
cnf(c_0_92, negated_conjecture, (r1_tarski(esk2_0,X1)|~r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_84, c_0_59])).
cnf(c_0_93, negated_conjecture, (k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)!=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_82])])).
cnf(c_0_94, plain, (X2=k3_enumset1(X1,X1,X1,X1,X1)|esk5_2(X1,X2)!=X1|~r2_hidden(esk5_2(X1,X2),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_33]), c_0_21]), c_0_30]), c_0_31])).
cnf(c_0_95, plain, (esk5_2(X1,k1_xboole_0)=X1), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_96, negated_conjecture, (esk6_1(esk3_0)=esk1_0|esk3_0=esk2_0), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_97, negated_conjecture, (~r2_hidden(esk1_0,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_93])).
cnf(c_0_98, plain, (r2_hidden(X1,k3_enumset1(X1,X1,X1,X1,X1))), inference(er,[status(thm)],[c_0_79])).
cnf(c_0_99, plain, (k3_enumset1(X1,X1,X1,X1,X1)=k1_xboole_0|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_100, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X2!=k2_xboole_0(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_101, negated_conjecture, (esk3_0=esk2_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_96]), c_0_97])).
cnf(c_0_102, plain, (~v1_xboole_0(k3_enumset1(X1,X1,X1,X1,X1))), inference(spm,[status(thm)],[c_0_72, c_0_98])).
cnf(c_0_103, plain, (k3_enumset1(X1,X1,X1,X1,X1)=esk2_0|~r2_hidden(X1,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_99, c_0_82]), c_0_82])).
cnf(c_0_104, plain, (v1_xboole_0(esk2_0)), inference(rw,[status(thm)],[c_0_88, c_0_82])).
cnf(c_0_105, plain, (r2_hidden(X1,X4)|r2_hidden(X1,X3)|X2!=k3_tarski(k3_enumset1(X3,X3,X3,X3,X4))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_100, c_0_29]), c_0_30]), c_0_31])).
cnf(c_0_106, negated_conjecture, (k3_tarski(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))=k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[c_0_38, c_0_101])).
cnf(c_0_107, plain, (~r2_hidden(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_104])])).
cnf(c_0_108, negated_conjecture, (X1!=k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)|~r2_hidden(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_107])).
cnf(c_0_109, negated_conjecture, (~r2_hidden(X1,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))), inference(er,[status(thm)],[c_0_108])).
cnf(c_0_110, negated_conjecture, (k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)!=esk2_0), inference(rw,[status(thm)],[c_0_93, c_0_101])).
cnf(c_0_111, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_90]), c_0_110]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 112
# Proof object clause steps            : 80
# Proof object formula steps           : 32
# Proof object conjectures             : 40
# Proof object clause conjectures      : 37
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 34
# Proof object simplifying inferences  : 94
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 25
# Processed clauses                    : 235
# ...of these trivial                  : 6
# ...subsumed                          : 67
# ...remaining for further processing  : 162
# Other redundant clauses eliminated   : 16
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 68
# Generated clauses                    : 642
# ...of the previous two non-trivial   : 607
# Contextual simplify-reflections      : 7
# Paramodulations                      : 601
# Factorizations                       : 6
# Equation resolutions                 : 35
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 63
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 48
# Current number of unprocessed clauses: 368
# ...number of literals in the above   : 1457
# Current number of archived formulas  : 0
# Current number of archived clauses   : 103
# Clause-clause subsumption calls (NU) : 1399
# Rec. Clause-clause subsumption calls : 1157
# Non-unit clause-clause subsumptions  : 64
# Unit Clause-clause subsumption calls : 146
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 9197

# -------------------------------------------------
# User time                : 0.019 s
# System time              : 0.000 s
# Total time               : 0.019 s
# Maximum resident set size: 28856 pages
