# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t15_yellow19, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2)))), file('data/MPTP/proofs/421bddf8d115426eb11c03b8c0a59eaa.E_input', t15_yellow19)).
fof(t2_yellow19, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(X1)))&(v2_waybel_0(X2,k3_yellow_1(X1))&(v13_waybel_0(X2,k3_yellow_1(X1))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X1))))))))=>![X3]:~((r2_hidden(X3,X2)&v1_xboole_0(X3))))), file('data/MPTP/proofs/421bddf8d115426eb11c03b8c0a59eaa.E_input', t2_yellow19)).
fof(d2_yellow_1, axiom, ![X1]:k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1)), file('data/MPTP/proofs/421bddf8d115426eb11c03b8c0a59eaa.E_input', d2_yellow_1)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/421bddf8d115426eb11c03b8c0a59eaa.E_input', d3_struct_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/421bddf8d115426eb11c03b8c0a59eaa.E_input', d1_xboole_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/421bddf8d115426eb11c03b8c0a59eaa.E_input', fc2_struct_0)).
fof(l24_zfmisc_1, axiom, ![X1, X2]:~((r1_xboole_0(k1_tarski(X1),X2)&r2_hidden(X1,X2))), file('data/MPTP/proofs/421bddf8d115426eb11c03b8c0a59eaa.E_input', l24_zfmisc_1)).
fof(t17_zfmisc_1, axiom, ![X1, X2]:(~(X1=X2)=>r1_xboole_0(k1_tarski(X1),k1_tarski(X2))), file('data/MPTP/proofs/421bddf8d115426eb11c03b8c0a59eaa.E_input', t17_zfmisc_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/421bddf8d115426eb11c03b8c0a59eaa.E_input', t3_xboole_0)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/421bddf8d115426eb11c03b8c0a59eaa.E_input', t65_xboole_1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/421bddf8d115426eb11c03b8c0a59eaa.E_input', t4_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/421bddf8d115426eb11c03b8c0a59eaa.E_input', d3_tarski)).
fof(t14_yellow19, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))))))))=>k7_subset_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))),X2,k1_tarski(k1_xboole_0))=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2)))), file('data/MPTP/proofs/421bddf8d115426eb11c03b8c0a59eaa.E_input', t14_yellow19)).
fof(redefinition_k7_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k7_subset_1(X1,X2,X3)=k4_xboole_0(X2,X3)), file('data/MPTP/proofs/421bddf8d115426eb11c03b8c0a59eaa.E_input', redefinition_k7_subset_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/421bddf8d115426eb11c03b8c0a59eaa.E_input', t100_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/421bddf8d115426eb11c03b8c0a59eaa.E_input', d10_xboole_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/421bddf8d115426eb11c03b8c0a59eaa.E_input', fc1_xboole_0)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/421bddf8d115426eb11c03b8c0a59eaa.E_input', t3_boole)).
fof(c_0_18, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2))))), inference(assume_negation,[status(cth)],[t15_yellow19])).
fof(c_0_19, plain, ![X6, X7, X8]:(v1_xboole_0(X6)|(v1_xboole_0(X7)|(~v1_subset_1(X7,u1_struct_0(k3_yellow_1(X6)))|(~v2_waybel_0(X7,k3_yellow_1(X6))|(~v13_waybel_0(X7,k3_yellow_1(X6))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X6)))))))|(~r2_hidden(X8,X7)|~v1_xboole_0(X8)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t2_yellow19])])])])).
fof(c_0_20, plain, ![X9]:k3_yellow_1(X9)=k3_lattice3(k1_lattice3(X9)), inference(variable_rename,[status(thm)],[d2_yellow_1])).
fof(c_0_21, plain, ![X10]:(~l1_struct_0(X10)|k2_struct_0(X10)=u1_struct_0(X10)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
fof(c_0_22, negated_conjecture, ((~v2_struct_0(esk1_0)&l1_struct_0(esk1_0))&((~v1_xboole_0(esk2_0)&(v1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))&(v2_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))&(v13_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))))))&esk2_0!=k2_yellow19(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_18])])])])).
cnf(c_0_23, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(X1)))|~v2_waybel_0(X2,k3_yellow_1(X1))|~v13_waybel_0(X2,k3_yellow_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X1))))|~r2_hidden(X3,X2)|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, plain, (k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_25, plain, ![X11, X12, X13]:((~v1_xboole_0(X11)|~r2_hidden(X12,X11))&(r2_hidden(esk3_1(X13),X13)|v1_xboole_0(X13))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_26, plain, ![X15]:(v2_struct_0(X15)|~l1_struct_0(X15)|~v1_xboole_0(u1_struct_0(X15))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_27, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, negated_conjecture, (l1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xboole_0(X3)|~r2_hidden(X3,X2)|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(X1)))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(X1)))|~v1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X1)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_24]), c_0_24]), c_0_24]), c_0_24])).
cnf(c_0_30, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32, negated_conjecture, (v13_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33, negated_conjecture, (v2_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_34, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_35, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, negated_conjecture, (u1_struct_0(esk1_0)=k2_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_37, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_38, plain, ![X16, X17]:(~r1_xboole_0(k1_tarski(X16),X17)|~r2_hidden(X16,X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l24_zfmisc_1])])).
fof(c_0_39, plain, ![X18, X19]:(X18=X19|r1_xboole_0(k1_tarski(X18),k1_tarski(X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_zfmisc_1])])])).
cnf(c_0_40, plain, (v1_xboole_0(X1)|~r2_hidden(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X1)))))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X1)))|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X1)))|~v1_subset_1(X3,u1_struct_0(k3_lattice3(k1_lattice3(X1))))|~v1_xboole_0(X2)), inference(csr,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_41, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))))), inference(rw,[status(thm)],[c_0_31, c_0_24])).
cnf(c_0_42, negated_conjecture, (v13_waybel_0(esk2_0,k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))), inference(rw,[status(thm)],[c_0_32, c_0_24])).
cnf(c_0_43, negated_conjecture, (v2_waybel_0(esk2_0,k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))), inference(rw,[status(thm)],[c_0_33, c_0_24])).
cnf(c_0_44, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(esk1_0)))))), inference(rw,[status(thm)],[c_0_34, c_0_24])).
cnf(c_0_45, negated_conjecture, (~v1_xboole_0(k2_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_28])]), c_0_37])).
fof(c_0_46, plain, ![X20, X21, X23, X24, X25]:(((r2_hidden(esk4_2(X20,X21),X20)|r1_xboole_0(X20,X21))&(r2_hidden(esk4_2(X20,X21),X21)|r1_xboole_0(X20,X21)))&(~r2_hidden(X25,X23)|~r2_hidden(X25,X24)|~r1_xboole_0(X23,X24))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_47, plain, (~r1_xboole_0(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_48, plain, (X1=X2|r1_xboole_0(k1_tarski(X1),k1_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_49, plain, ![X26]:r1_xboole_0(X26,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
fof(c_0_50, plain, ![X27, X28, X30, X31, X32]:((r1_xboole_0(X27,X28)|r2_hidden(esk5_2(X27,X28),k3_xboole_0(X27,X28)))&(~r2_hidden(X32,k3_xboole_0(X30,X31))|~r1_xboole_0(X30,X31))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
cnf(c_0_51, negated_conjecture, (~r2_hidden(X1,esk2_0)|~v1_xboole_0(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42]), c_0_43]), c_0_44])]), c_0_45])).
cnf(c_0_52, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_53, plain, (X1=X2|~r2_hidden(X1,k1_tarski(X2))), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_54, plain, (r2_hidden(esk4_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_55, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_56, plain, ![X33, X34, X35, X36, X37]:((~r1_tarski(X33,X34)|(~r2_hidden(X35,X33)|r2_hidden(X35,X34)))&((r2_hidden(esk6_2(X36,X37),X36)|r1_tarski(X36,X37))&(~r2_hidden(esk6_2(X36,X37),X37)|r1_tarski(X36,X37)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_57, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_58, plain, (r2_hidden(esk3_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_59, plain, ![X39, X40]:(v2_struct_0(X39)|~l1_struct_0(X39)|(v1_xboole_0(X40)|(~v2_waybel_0(X40,k3_yellow_1(k2_struct_0(X39)))|(~v13_waybel_0(X40,k3_yellow_1(k2_struct_0(X39)))|~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X39)))))))|k7_subset_1(u1_struct_0(k3_yellow_1(k2_struct_0(X39))),X40,k1_tarski(k1_xboole_0))=k2_yellow19(X39,k3_yellow19(X39,k2_struct_0(X39),X40)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_yellow19])])])])).
fof(c_0_60, plain, ![X41, X42, X43]:(~m1_subset_1(X42,k1_zfmisc_1(X41))|k7_subset_1(X41,X42,X43)=k4_xboole_0(X42,X43)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_subset_1])])).
fof(c_0_61, plain, ![X44, X45]:k4_xboole_0(X44,X45)=k5_xboole_0(X44,k3_xboole_0(X44,X45)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_62, negated_conjecture, (r1_xboole_0(esk2_0,X1)|~v1_xboole_0(esk4_2(esk2_0,X1))), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_63, plain, (esk4_2(X1,k1_tarski(X2))=X2|r1_xboole_0(X1,k1_tarski(X2))), inference(spm,[status(thm)],[c_0_53, c_0_54])).
fof(c_0_64, plain, ![X46, X47]:(((r1_tarski(X46,X47)|X46!=X47)&(r1_tarski(X47,X46)|X46!=X47))&(~r1_tarski(X46,X47)|~r1_tarski(X47,X46)|X46=X47)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_65, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_47, c_0_55])).
cnf(c_0_66, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_67, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_68, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|k7_subset_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))),X2,k1_tarski(k1_xboole_0))=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2))|~l1_struct_0(X1)|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_69, plain, (k7_subset_1(X2,X1,X3)=k4_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_70, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_71, negated_conjecture, (r1_xboole_0(esk2_0,k1_tarski(X1))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_72, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
fof(c_0_73, plain, ![X48]:k4_xboole_0(X48,k1_xboole_0)=X48, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_74, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_75, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_76, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_30, c_0_66])).
cnf(c_0_77, plain, (v1_xboole_0(k3_xboole_0(X1,k1_xboole_0))), inference(spm,[status(thm)],[c_0_67, c_0_55])).
cnf(c_0_78, plain, (k7_subset_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(X1)))),X2,k1_tarski(k1_xboole_0))=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2))|v2_struct_0(X1)|v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(X1))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_24]), c_0_24]), c_0_24]), c_0_24])).
cnf(c_0_79, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_80, plain, (k7_subset_1(X2,X1,X3)=k5_xboole_0(X1,k3_xboole_0(X1,X3))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_81, plain, (r1_tarski(k3_xboole_0(X1,X2),X3)|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_57, c_0_66])).
cnf(c_0_82, negated_conjecture, (r1_xboole_0(esk2_0,k1_tarski(k1_xboole_0))), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_83, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_84, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_85, plain, (r1_tarski(k3_xboole_0(X1,k1_xboole_0),X2)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_86, negated_conjecture, (k7_subset_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(esk1_0)))),esk2_0,k1_tarski(k1_xboole_0))=k2_yellow19(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_41]), c_0_42]), c_0_43]), c_0_28])]), c_0_79]), c_0_37])).
cnf(c_0_87, negated_conjecture, (k7_subset_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(esk1_0)))),esk2_0,X1)=k5_xboole_0(esk2_0,k3_xboole_0(esk2_0,X1))), inference(spm,[status(thm)],[c_0_80, c_0_41])).
cnf(c_0_88, negated_conjecture, (r1_tarski(k3_xboole_0(esk2_0,k1_tarski(k1_xboole_0)),X1)), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_89, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[c_0_83, c_0_70])).
cnf(c_0_90, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_91, negated_conjecture, (k5_xboole_0(esk2_0,k3_xboole_0(esk2_0,k1_tarski(k1_xboole_0)))=k2_yellow19(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0))), inference(rw,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_92, negated_conjecture, (k3_xboole_0(esk2_0,k1_tarski(k1_xboole_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_84, c_0_88])).
cnf(c_0_93, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_94, negated_conjecture, (esk2_0!=k2_yellow19(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_95, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_92]), c_0_93]), c_0_94]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 96
# Proof object clause steps            : 60
# Proof object formula steps           : 36
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 34
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 31
# Processed clauses                    : 139
# ...of these trivial                  : 4
# ...subsumed                          : 15
# ...remaining for further processing  : 120
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 17
# Generated clauses                    : 182
# ...of the previous two non-trivial   : 147
# Contextual simplify-reflections      : 1
# Paramodulations                      : 180
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 71
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 48
# Current number of unprocessed clauses: 48
# ...number of literals in the above   : 99
# Current number of archived formulas  : 0
# Current number of archived clauses   : 49
# Clause-clause subsumption calls (NU) : 479
# Rec. Clause-clause subsumption calls : 211
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 76
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4201

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28680 pages
