# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', t70_enumset1)).
fof(t14_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&(r1_tarski(X3,X2)&![X4]:((r1_tarski(X1,X4)&r1_tarski(X3,X4))=>r1_tarski(X2,X4))))=>X2=k2_xboole_0(X1,X3)), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', t14_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', d10_xboole_0)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', t3_xboole_0)).
fof(t7_tarski, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&r2_hidden(X4,X3))))))), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', t7_tarski)).
fof(t10_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(X1,k2_xboole_0(X3,X2))), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', t10_xboole_1)).
fof(t31_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>r1_tarski(k3_relat_1(X1),k3_relat_1(X2))))), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', t31_relat_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', t7_xboole_1)).
fof(t9_tarski, axiom, ![X1]:?[X2]:(r2_hidden(X1,X2)&(![X3, X4]:((r2_hidden(X3,X2)&r1_tarski(X4,X3))=>r2_hidden(X4,X2))&(![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&![X5]:(r1_tarski(X5,X3)=>r2_hidden(X5,X4))))))&![X3]:~((r1_tarski(X3,X2)&(~(r2_tarski(X3,X2))&~(r2_hidden(X3,X2)))))))), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', t9_tarski)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', t65_xboole_1)).
fof(t8_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X3,X2))=>r1_tarski(k2_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', t8_xboole_1)).
fof(d6_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k3_relat_1(X1)=k2_xboole_0(k1_relat_1(X1),k2_relat_1(X1))), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', d6_relat_1)).
fof(t43_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3)), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', t43_xboole_1)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', d4_relat_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', commutativity_k2_tarski)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', t2_xboole_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', d5_relat_1)).
fof(t44_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k4_xboole_0(X1,X2),X3)=>r1_tarski(X1,k2_xboole_0(X2,X3))), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', t44_xboole_1)).
fof(t15_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>r1_tarski(k6_subset_1(k1_relat_1(X1),k1_relat_1(X2)),k1_relat_1(k6_subset_1(X1,X2))))), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', t15_relat_1)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', redefinition_k6_subset_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', t7_xboole_0)).
fof(t28_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>r1_tarski(k6_subset_1(k2_relat_1(X1),k2_relat_1(X2)),k2_relat_1(k6_subset_1(X1,X2))))), file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', t28_relat_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/4232cac600874bd2b6a6470115c6e09b.E_input', t3_boole)).
fof(c_0_24, plain, ![X6, X7]:k3_tarski(k2_tarski(X6,X7))=k2_xboole_0(X6,X7), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_25, plain, ![X8, X9]:k1_enumset1(X8,X8,X9)=k2_tarski(X8,X9), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_26, plain, ![X22, X23, X24]:(((r1_tarski(X22,esk1_3(X22,X23,X24))|~r1_tarski(X24,X23)|~r1_tarski(X22,X23)|X23=k2_xboole_0(X22,X24))&(r1_tarski(X24,esk1_3(X22,X23,X24))|~r1_tarski(X24,X23)|~r1_tarski(X22,X23)|X23=k2_xboole_0(X22,X24)))&(~r1_tarski(X23,esk1_3(X22,X23,X24))|~r1_tarski(X24,X23)|~r1_tarski(X22,X23)|X23=k2_xboole_0(X22,X24))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t14_xboole_1])])])])).
cnf(c_0_27, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_28, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_29, plain, ![X15, X16]:(((r1_tarski(X15,X16)|X15!=X16)&(r1_tarski(X16,X15)|X15!=X16))&(~r1_tarski(X15,X16)|~r1_tarski(X16,X15)|X15=X16)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_30, plain, ![X32, X33, X35, X36, X37]:(((r2_hidden(esk4_2(X32,X33),X32)|r1_xboole_0(X32,X33))&(r2_hidden(esk4_2(X32,X33),X33)|r1_xboole_0(X32,X33)))&(~r2_hidden(X37,X35)|~r2_hidden(X37,X36)|~r1_xboole_0(X35,X36))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_31, plain, ![X39, X40, X42]:((r2_hidden(esk5_2(X39,X40),X40)|~r2_hidden(X39,X40))&(~r2_hidden(X42,X40)|~r2_hidden(X42,esk5_2(X39,X40))|~r2_hidden(X39,X40))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tarski])])])])])).
fof(c_0_32, plain, ![X77, X78, X79]:(~r1_tarski(X77,X78)|r1_tarski(X77,k2_xboole_0(X79,X78))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_xboole_1])])).
cnf(c_0_33, plain, (X1=k2_xboole_0(X2,X3)|~r1_tarski(X1,esk1_3(X2,X1,X3))|~r1_tarski(X3,X1)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_35, plain, (r1_tarski(X1,esk1_3(X1,X2,X3))|X2=k2_xboole_0(X1,X3)|~r1_tarski(X3,X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, plain, (r2_hidden(esk5_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_39, plain, (r1_tarski(X1,k2_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40, plain, (X1=k3_tarski(k1_enumset1(X2,X2,X3))|~r1_tarski(X3,X1)|~r1_tarski(X2,X1)|~r1_tarski(X1,esk1_3(X2,X1,X3))), inference(rw,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_41, plain, (X2=k3_tarski(k1_enumset1(X1,X1,X3))|r1_tarski(X1,esk1_3(X1,X2,X3))|~r1_tarski(X3,X2)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_35, c_0_34])).
cnf(c_0_42, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_36])).
fof(c_0_43, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>r1_tarski(k3_relat_1(X1),k3_relat_1(X2)))))), inference(assume_negation,[status(cth)],[t31_relat_1])).
cnf(c_0_44, plain, (~r2_hidden(esk5_2(X1,X2),X3)|~r2_hidden(X1,X2)|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_45, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X3,X3,X2)))|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_39, c_0_34])).
cnf(c_0_46, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=X1|~r1_tarski(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42])])).
fof(c_0_47, negated_conjecture, (v1_relat_1(esk2_0)&(v1_relat_1(esk3_0)&(r1_tarski(esk2_0,esk3_0)&~r1_tarski(k3_relat_1(esk2_0),k3_relat_1(esk3_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_43])])])).
fof(c_0_48, plain, ![X17, X18]:r1_tarski(X17,k2_xboole_0(X17,X18)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_49, plain, ![X43, X45, X46, X47, X49, X50]:(r2_hidden(X43,esk6_1(X43))&((~r2_hidden(X45,esk6_1(X43))|~r1_tarski(X46,X45)|r2_hidden(X46,esk6_1(X43)))&(((r2_hidden(esk7_2(X43,X47),esk6_1(X43))|~r2_hidden(X47,esk6_1(X43)))&(~r1_tarski(X49,X47)|r2_hidden(X49,esk7_2(X43,X47))|~r2_hidden(X47,esk6_1(X43))))&(~r1_tarski(X50,esk6_1(X43))|(r2_tarski(X50,esk6_1(X43))|r2_hidden(X50,esk6_1(X43))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_tarski])])])])])])).
cnf(c_0_50, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_44, c_0_38])).
cnf(c_0_51, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_52, plain, ![X38]:r1_xboole_0(X38,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
fof(c_0_53, plain, ![X62, X63, X64]:(~r1_tarski(X62,X63)|~r1_tarski(X64,X63)|r1_tarski(k2_xboole_0(X62,X64),X63)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_xboole_1])])).
fof(c_0_54, plain, ![X65]:(~v1_relat_1(X65)|k3_relat_1(X65)=k2_xboole_0(k1_relat_1(X65),k2_relat_1(X65))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_relat_1])])).
fof(c_0_55, plain, ![X10, X11, X12]:(~r1_tarski(X10,k2_xboole_0(X11,X12))|r1_tarski(k4_xboole_0(X10,X11),X12)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
cnf(c_0_56, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,X3)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_57, negated_conjecture, (r1_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_58, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59, plain, (r2_hidden(X1,esk6_1(X1))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_60, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_61, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_52])).
fof(c_0_62, plain, ![X51, X52, X53, X55, X56, X57, X58, X60]:(((~r2_hidden(X53,X52)|r2_hidden(k4_tarski(X53,esk8_3(X51,X52,X53)),X51)|X52!=k1_relat_1(X51))&(~r2_hidden(k4_tarski(X55,X56),X51)|r2_hidden(X55,X52)|X52!=k1_relat_1(X51)))&((~r2_hidden(esk9_2(X57,X58),X58)|~r2_hidden(k4_tarski(esk9_2(X57,X58),X60),X57)|X58=k1_relat_1(X57))&(r2_hidden(esk9_2(X57,X58),X58)|r2_hidden(k4_tarski(esk9_2(X57,X58),esk10_2(X57,X58)),X57)|X58=k1_relat_1(X57)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
cnf(c_0_63, plain, (r1_tarski(k2_xboole_0(X1,X3),X2)|~r1_tarski(X1,X2)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_64, plain, (k3_relat_1(X1)=k2_xboole_0(k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
fof(c_0_65, plain, ![X19, X20]:k2_tarski(X19,X20)=k2_tarski(X20,X19), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
fof(c_0_66, plain, ![X31]:r1_tarski(k1_xboole_0,X31), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_67, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_68, negated_conjecture, (r1_tarski(esk2_0,X1)|~r1_tarski(esk3_0,X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_69, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[c_0_58, c_0_34])).
cnf(c_0_70, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(X2,esk6_1(X1))), inference(spm,[status(thm)],[c_0_37, c_0_59])).
cnf(c_0_71, plain, (r1_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_72, plain, (r2_hidden(k4_tarski(X1,esk8_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_62])).
fof(c_0_73, plain, ![X66, X67, X68, X70, X71, X72, X73, X75]:(((~r2_hidden(X68,X67)|r2_hidden(k4_tarski(esk11_3(X66,X67,X68),X68),X66)|X67!=k2_relat_1(X66))&(~r2_hidden(k4_tarski(X71,X70),X66)|r2_hidden(X70,X67)|X67!=k2_relat_1(X66)))&((~r2_hidden(esk12_2(X72,X73),X73)|~r2_hidden(k4_tarski(X75,esk12_2(X72,X73)),X72)|X73=k2_relat_1(X72))&(r2_hidden(esk12_2(X72,X73),X73)|r2_hidden(k4_tarski(esk13_2(X72,X73),esk12_2(X72,X73)),X72)|X73=k2_relat_1(X72)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_74, plain, (r1_tarski(k3_tarski(k1_enumset1(X1,X1,X3)),X2)|~r1_tarski(X3,X2)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_63, c_0_34])).
cnf(c_0_75, plain, (k3_relat_1(X1)=k3_tarski(k1_enumset1(k1_relat_1(X1),k1_relat_1(X1),k2_relat_1(X1)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_64, c_0_34])).
cnf(c_0_76, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
fof(c_0_77, plain, ![X28, X29, X30]:(~r1_tarski(k4_xboole_0(X28,X29),X30)|r1_tarski(X28,k2_xboole_0(X29,X30))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t44_xboole_1])])).
fof(c_0_78, plain, ![X80, X81]:(~v1_relat_1(X80)|(~v1_relat_1(X81)|r1_tarski(k6_subset_1(k1_relat_1(X80),k1_relat_1(X81)),k1_relat_1(k6_subset_1(X80,X81))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_relat_1])])])).
fof(c_0_79, plain, ![X13, X14]:k6_subset_1(X13,X14)=k4_xboole_0(X13,X14), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
cnf(c_0_80, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_81, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_82, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k3_tarski(k1_enumset1(X2,X2,X3)))), inference(rw,[status(thm)],[c_0_67, c_0_34])).
cnf(c_0_83, negated_conjecture, (r1_tarski(esk2_0,k3_tarski(k1_enumset1(esk3_0,esk3_0,X1)))), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_84, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_85, plain, (r2_hidden(k4_tarski(X1,esk8_3(X2,k1_relat_1(X2),X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_72])).
fof(c_0_86, plain, ![X82]:(X82=k1_xboole_0|r2_hidden(esk14_1(X82),X82)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_87, plain, (r2_hidden(k4_tarski(esk11_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_88, negated_conjecture, (~r1_tarski(k3_relat_1(esk2_0),k3_relat_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_89, plain, (r1_tarski(k3_relat_1(X1),X2)|~v1_relat_1(X1)|~r1_tarski(k2_relat_1(X1),X2)|~r1_tarski(k1_relat_1(X1),X2)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_90, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_91, plain, (k1_enumset1(X1,X1,X2)=k1_enumset1(X2,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_28]), c_0_28])).
cnf(c_0_92, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(k4_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_93, plain, (r1_tarski(k6_subset_1(k1_relat_1(X1),k1_relat_1(X2)),k1_relat_1(k6_subset_1(X1,X2)))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_94, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_95, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_96, negated_conjecture, (r1_tarski(k4_xboole_0(esk2_0,esk3_0),X1)), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_97, plain, (~r2_hidden(X1,k1_relat_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_98, plain, (X1=k1_xboole_0|r2_hidden(esk14_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
fof(c_0_99, plain, ![X84, X85]:(~v1_relat_1(X84)|(~v1_relat_1(X85)|r1_tarski(k6_subset_1(k2_relat_1(X84),k2_relat_1(X85)),k2_relat_1(k6_subset_1(X84,X85))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_relat_1])])])).
cnf(c_0_100, plain, (r2_hidden(k4_tarski(esk11_3(X1,k2_relat_1(X1),X2),X2),X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_87])).
cnf(c_0_101, negated_conjecture, (~r1_tarski(k2_relat_1(esk2_0),k3_relat_1(esk3_0))|~r1_tarski(k1_relat_1(esk2_0),k3_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_90])])).
cnf(c_0_102, plain, (r1_tarski(X1,k3_relat_1(X2))|~v1_relat_1(X2)|~r1_tarski(X1,k2_relat_1(X2))), inference(spm,[status(thm)],[c_0_45, c_0_75])).
cnf(c_0_103, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_104, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X2,X2,X3)))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_45, c_0_91])).
cnf(c_0_105, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k3_tarski(k1_enumset1(X3,X3,X2)))), inference(spm,[status(thm)],[c_0_82, c_0_91])).
cnf(c_0_106, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X2,X2,X3)))|~r1_tarski(k4_xboole_0(X1,X2),X3)), inference(rw,[status(thm)],[c_0_92, c_0_34])).
cnf(c_0_107, plain, (r1_tarski(k4_xboole_0(k1_relat_1(X1),k1_relat_1(X2)),k1_relat_1(k4_xboole_0(X1,X2)))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_94]), c_0_94])).
cnf(c_0_108, negated_conjecture, (k4_xboole_0(esk2_0,esk3_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_109, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_97, c_0_98])).
fof(c_0_110, plain, ![X21]:k4_xboole_0(X21,k1_xboole_0)=X21, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_111, plain, (r1_tarski(k6_subset_1(k2_relat_1(X1),k2_relat_1(X2)),k2_relat_1(k6_subset_1(X1,X2)))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_112, plain, (~r2_hidden(X1,k2_relat_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_84, c_0_100])).
cnf(c_0_113, negated_conjecture, (~r1_tarski(k1_relat_1(esk2_0),k3_relat_1(esk3_0))|~r1_tarski(k2_relat_1(esk2_0),k2_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_103])])).
cnf(c_0_114, plain, (r1_tarski(X1,k3_relat_1(X2))|~v1_relat_1(X2)|~r1_tarski(X1,k1_relat_1(X2))), inference(spm,[status(thm)],[c_0_104, c_0_75])).
cnf(c_0_115, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(k4_xboole_0(X1,X3),X2)), inference(spm,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_116, negated_conjecture, (r1_tarski(k4_xboole_0(k1_relat_1(esk2_0),k1_relat_1(esk3_0)),k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_109]), c_0_103]), c_0_90])])).
cnf(c_0_117, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_118, plain, (r1_tarski(k4_xboole_0(k2_relat_1(X1),k2_relat_1(X2)),k2_relat_1(k4_xboole_0(X1,X2)))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_111, c_0_94]), c_0_94])).
cnf(c_0_119, plain, (k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_112, c_0_98])).
cnf(c_0_120, negated_conjecture, (~r1_tarski(k2_relat_1(esk2_0),k2_relat_1(esk3_0))|~r1_tarski(k1_relat_1(esk2_0),k1_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_103])])).
cnf(c_0_121, negated_conjecture, (r1_tarski(k1_relat_1(esk2_0),k1_relat_1(esk3_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_117])).
cnf(c_0_122, negated_conjecture, (r1_tarski(k4_xboole_0(k2_relat_1(esk2_0),k2_relat_1(esk3_0)),k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_108]), c_0_119]), c_0_103]), c_0_90])])).
cnf(c_0_123, negated_conjecture, (~r1_tarski(k2_relat_1(esk2_0),k2_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_120, c_0_121])])).
cnf(c_0_124, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_122]), c_0_117]), c_0_123]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 125
# Proof object clause steps            : 76
# Proof object formula steps           : 49
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 30
# Proof object initial formulas used   : 24
# Proof object generating inferences   : 30
# Proof object simplifying inferences  : 39
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 24
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 44
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 41
# Processed clauses                    : 1288
# ...of these trivial                  : 6
# ...subsumed                          : 761
# ...remaining for further processing  : 521
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 8
# Backward-rewritten                   : 19
# Generated clauses                    : 4211
# ...of the previous two non-trivial   : 3594
# Contextual simplify-reflections      : 0
# Paramodulations                      : 4205
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 448
#    Positive orientable unit clauses  : 39
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 11
#    Non-unit-clauses                  : 397
# Current number of unprocessed clauses: 2315
# ...number of literals in the above   : 6873
# Current number of archived formulas  : 0
# Current number of archived clauses   : 70
# Clause-clause subsumption calls (NU) : 23688
# Rec. Clause-clause subsumption calls : 12264
# Non-unit clause-clause subsumptions  : 632
# Unit Clause-clause subsumption calls : 607
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 79
# BW rewrite match successes           : 19
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 52247

# -------------------------------------------------
# User time                : 0.048 s
# System time              : 0.005 s
# Total time               : 0.052 s
# Maximum resident set size: 28856 pages
