# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/424884421f8e4bb296d4594d45224c0a.E_input', d9_yellow_0)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/424884421f8e4bb296d4594d45224c0a.E_input', dt_k1_yellow_0)).
fof(t5_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r2_lattice3(X1,X2,X3)=>r2_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3))&((r2_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3)=>r2_lattice3(X1,X2,X3))&((r1_lattice3(X1,X2,X3)=>r1_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3))&(r1_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3)=>r1_lattice3(X1,X2,X3))))))), file('data/MPTP/proofs/424884421f8e4bb296d4594d45224c0a.E_input', t5_yellow_0)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/424884421f8e4bb296d4594d45224c0a.E_input', t12_setfam_1)).
fof(t50_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((r1_yellow_0(X1,X2)=>r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))&((r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))=>r1_yellow_0(X1,X2))&((r2_yellow_0(X1,X2)=>r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))&(r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))=>r2_yellow_0(X1,X2)))))), file('data/MPTP/proofs/424884421f8e4bb296d4594d45224c0a.E_input', t50_yellow_0)).
fof(t51_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((r1_yellow_0(X1,X2)|r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))=>k1_yellow_0(X1,X2)=k1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))), file('data/MPTP/proofs/424884421f8e4bb296d4594d45224c0a.E_input', t51_yellow_0)).
fof(c_0_6, plain, ![X11, X12, X13, X14]:(((r2_lattice3(X11,X12,X13)|X13!=k1_yellow_0(X11,X12)|~r1_yellow_0(X11,X12)|~m1_subset_1(X13,u1_struct_0(X11))|~l1_orders_2(X11))&(~m1_subset_1(X14,u1_struct_0(X11))|(~r2_lattice3(X11,X12,X14)|r1_orders_2(X11,X13,X14))|X13!=k1_yellow_0(X11,X12)|~r1_yellow_0(X11,X12)|~m1_subset_1(X13,u1_struct_0(X11))|~l1_orders_2(X11)))&((m1_subset_1(esk3_3(X11,X12,X13),u1_struct_0(X11))|~r2_lattice3(X11,X12,X13)|X13=k1_yellow_0(X11,X12)|~r1_yellow_0(X11,X12)|~m1_subset_1(X13,u1_struct_0(X11))|~l1_orders_2(X11))&((r2_lattice3(X11,X12,esk3_3(X11,X12,X13))|~r2_lattice3(X11,X12,X13)|X13=k1_yellow_0(X11,X12)|~r1_yellow_0(X11,X12)|~m1_subset_1(X13,u1_struct_0(X11))|~l1_orders_2(X11))&(~r1_orders_2(X11,X13,esk3_3(X11,X12,X13))|~r2_lattice3(X11,X12,X13)|X13=k1_yellow_0(X11,X12)|~r1_yellow_0(X11,X12)|~m1_subset_1(X13,u1_struct_0(X11))|~l1_orders_2(X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
fof(c_0_7, plain, ![X9, X10]:(~l1_orders_2(X9)|m1_subset_1(k1_yellow_0(X9,X10),u1_struct_0(X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
cnf(c_0_8, plain, (r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~r1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_10, plain, ![X18, X19, X20]:((~r2_lattice3(X18,X19,X20)|r2_lattice3(X18,k3_xboole_0(X19,u1_struct_0(X18)),X20)|~m1_subset_1(X20,u1_struct_0(X18))|(v2_struct_0(X18)|~l1_orders_2(X18)))&((~r2_lattice3(X18,k3_xboole_0(X19,u1_struct_0(X18)),X20)|r2_lattice3(X18,X19,X20)|~m1_subset_1(X20,u1_struct_0(X18))|(v2_struct_0(X18)|~l1_orders_2(X18)))&((~r1_lattice3(X18,X19,X20)|r1_lattice3(X18,k3_xboole_0(X19,u1_struct_0(X18)),X20)|~m1_subset_1(X20,u1_struct_0(X18))|(v2_struct_0(X18)|~l1_orders_2(X18)))&(~r1_lattice3(X18,k3_xboole_0(X19,u1_struct_0(X18)),X20)|r1_lattice3(X18,X19,X20)|~m1_subset_1(X20,u1_struct_0(X18))|(v2_struct_0(X18)|~l1_orders_2(X18)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t5_yellow_0])])])])])).
fof(c_0_11, plain, ![X16, X17]:k1_setfam_1(k2_tarski(X16,X17))=k3_xboole_0(X16,X17), inference(variable_rename,[status(thm)],[t12_setfam_1])).
cnf(c_0_12, plain, (X2=k1_yellow_0(X1,X3)|~r1_orders_2(X1,X2,esk3_3(X1,X3,X2))|~r2_lattice3(X1,X3,X2)|~r1_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_8]), c_0_9])).
cnf(c_0_14, plain, (r2_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3)|v2_struct_0(X1)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_16, plain, ![X7, X8]:((~r1_yellow_0(X7,X8)|r1_yellow_0(X7,k3_xboole_0(X8,u1_struct_0(X7)))|(v2_struct_0(X7)|~l1_orders_2(X7)))&((~r1_yellow_0(X7,k3_xboole_0(X8,u1_struct_0(X7)))|r1_yellow_0(X7,X8)|(v2_struct_0(X7)|~l1_orders_2(X7)))&((~r2_yellow_0(X7,X8)|r2_yellow_0(X7,k3_xboole_0(X8,u1_struct_0(X7)))|(v2_struct_0(X7)|~l1_orders_2(X7)))&(~r2_yellow_0(X7,k3_xboole_0(X8,u1_struct_0(X7)))|r2_yellow_0(X7,X8)|(v2_struct_0(X7)|~l1_orders_2(X7)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_yellow_0])])])])])).
cnf(c_0_17, plain, (r2_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r2_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, plain, (r2_lattice3(X1,X2,X3)|X3!=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_19, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((r1_yellow_0(X1,X2)|r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))=>k1_yellow_0(X1,X2)=k1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))))), inference(assume_negation,[status(cth)],[t51_yellow_0])).
cnf(c_0_20, plain, (k1_yellow_0(X1,X2)=k1_yellow_0(X1,X3)|~r2_lattice3(X1,X2,esk3_3(X1,X3,k1_yellow_0(X1,X2)))|~r2_lattice3(X1,X3,k1_yellow_0(X1,X2))|~m1_subset_1(esk3_3(X1,X3,k1_yellow_0(X1,X2)),u1_struct_0(X1))|~r1_yellow_0(X1,X3)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_9])).
cnf(c_0_21, plain, (v2_struct_0(X1)|r2_lattice3(X1,k1_setfam_1(k2_tarski(X2,u1_struct_0(X1))),X3)|~l1_orders_2(X1)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(rw,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_22, plain, (r1_yellow_0(X1,X2)|v2_struct_0(X1)|~r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (v2_struct_0(X1)|r2_lattice3(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_lattice3(X1,k1_setfam_1(k2_tarski(X2,u1_struct_0(X1))),X3)), inference(rw,[status(thm)],[c_0_17, c_0_15])).
cnf(c_0_24, plain, (r2_lattice3(X1,X2,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_18]), c_0_9])).
fof(c_0_25, negated_conjecture, ((~v2_struct_0(esk1_0)&l1_orders_2(esk1_0))&((r1_yellow_0(esk1_0,esk2_0)|r1_yellow_0(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0))))&k1_yellow_0(esk1_0,esk2_0)!=k1_yellow_0(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_19])])])])).
cnf(c_0_26, plain, (k1_yellow_0(X1,k1_setfam_1(k2_tarski(X2,u1_struct_0(X1))))=k1_yellow_0(X1,X3)|v2_struct_0(X1)|~r2_lattice3(X1,X2,esk3_3(X1,X3,k1_yellow_0(X1,k1_setfam_1(k2_tarski(X2,u1_struct_0(X1))))))|~r2_lattice3(X1,X3,k1_yellow_0(X1,k1_setfam_1(k2_tarski(X2,u1_struct_0(X1)))))|~m1_subset_1(esk3_3(X1,X3,k1_yellow_0(X1,k1_setfam_1(k2_tarski(X2,u1_struct_0(X1))))),u1_struct_0(X1))|~r1_yellow_0(X1,k1_setfam_1(k2_tarski(X2,u1_struct_0(X1))))|~r1_yellow_0(X1,X3)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27, plain, (r2_lattice3(X1,X2,esk3_3(X1,X2,X3))|X3=k1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_28, plain, (v2_struct_0(X1)|r1_yellow_0(X1,X2)|~l1_orders_2(X1)|~r1_yellow_0(X1,k1_setfam_1(k2_tarski(X2,u1_struct_0(X1))))), inference(rw,[status(thm)],[c_0_22, c_0_15])).
cnf(c_0_29, plain, (r2_lattice3(X1,X2,k1_yellow_0(X1,k1_setfam_1(k2_tarski(X2,u1_struct_0(X1)))))|v2_struct_0(X1)|~r1_yellow_0(X1,k1_setfam_1(k2_tarski(X2,u1_struct_0(X1))))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_9])).
cnf(c_0_30, negated_conjecture, (k1_yellow_0(esk1_0,esk2_0)!=k1_yellow_0(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, plain, (k1_yellow_0(X1,k1_setfam_1(k2_tarski(X2,u1_struct_0(X1))))=k1_yellow_0(X1,X2)|v2_struct_0(X1)|~m1_subset_1(esk3_3(X1,X2,k1_yellow_0(X1,k1_setfam_1(k2_tarski(X2,u1_struct_0(X1))))),u1_struct_0(X1))|~r1_yellow_0(X1,k1_setfam_1(k2_tarski(X2,u1_struct_0(X1))))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_9]), c_0_29])).
cnf(c_0_32, plain, (m1_subset_1(esk3_3(X1,X2,X3),u1_struct_0(X1))|X3=k1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_33, negated_conjecture, (r1_yellow_0(esk1_0,esk2_0)|r1_yellow_0(esk1_0,k3_xboole_0(esk2_0,u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, negated_conjecture, (k1_yellow_0(esk1_0,esk2_0)!=k1_yellow_0(esk1_0,k1_setfam_1(k2_tarski(esk2_0,u1_struct_0(esk1_0))))), inference(rw,[status(thm)],[c_0_30, c_0_15])).
cnf(c_0_35, plain, (k1_yellow_0(X1,k1_setfam_1(k2_tarski(X2,u1_struct_0(X1))))=k1_yellow_0(X1,X2)|v2_struct_0(X1)|~r1_yellow_0(X1,k1_setfam_1(k2_tarski(X2,u1_struct_0(X1))))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_28]), c_0_9]), c_0_29])).
cnf(c_0_36, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_37, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_38, plain, (r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))|v2_struct_0(X1)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_39, negated_conjecture, (r1_yellow_0(esk1_0,esk2_0)|r1_yellow_0(esk1_0,k1_setfam_1(k2_tarski(esk2_0,u1_struct_0(esk1_0))))), inference(rw,[status(thm)],[c_0_33, c_0_15])).
cnf(c_0_40, negated_conjecture, (~r1_yellow_0(esk1_0,k1_setfam_1(k2_tarski(esk2_0,u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])]), c_0_37])).
cnf(c_0_41, plain, (v2_struct_0(X1)|r1_yellow_0(X1,k1_setfam_1(k2_tarski(X2,u1_struct_0(X1))))|~l1_orders_2(X1)|~r1_yellow_0(X1,X2)), inference(rw,[status(thm)],[c_0_38, c_0_15])).
cnf(c_0_42, negated_conjecture, (r1_yellow_0(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_39]), c_0_36])]), c_0_37])).
cnf(c_0_43, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42]), c_0_36])]), c_0_37]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 44
# Proof object clause steps            : 31
# Proof object formula steps           : 13
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 28
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 18
# Processed clauses                    : 46
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 46
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 1
# Generated clauses                    : 25
# ...of the previous two non-trivial   : 20
# Contextual simplify-reflections      : 11
# Paramodulations                      : 23
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 24
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 19
# Current number of unprocessed clauses: 10
# ...number of literals in the above   : 59
# Current number of archived formulas  : 0
# Current number of archived clauses   : 21
# Clause-clause subsumption calls (NU) : 406
# Rec. Clause-clause subsumption calls : 28
# Non-unit clause-clause subsumptions  : 12
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2752

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.003 s
# Total time               : 0.011 s
# Maximum resident set size: 28584 pages
