# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t58_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k1_enumset1(X1,X2,X3),k2_enumset1(X4,X5,X6,X7)), file('data/MPTP/proofs/42a5356dbbf6454093949020b48842d0.E_input', t58_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/42a5356dbbf6454093949020b48842d0.E_input', t71_enumset1)).
fof(t67_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k4_enumset1(X1,X2,X3,X4,X5,X6),k2_tarski(X7,X8)), file('data/MPTP/proofs/42a5356dbbf6454093949020b48842d0.E_input', t67_enumset1)).
fof(t83_enumset1, axiom, ![X1, X2]:k3_enumset1(X1,X1,X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/42a5356dbbf6454093949020b48842d0.E_input', t83_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/42a5356dbbf6454093949020b48842d0.E_input', t73_enumset1)).
fof(l75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8)), file('data/MPTP/proofs/42a5356dbbf6454093949020b48842d0.E_input', l75_enumset1)).
fof(t91_enumset1, axiom, ![X1]:k4_enumset1(X1,X1,X1,X1,X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/42a5356dbbf6454093949020b48842d0.E_input', t91_enumset1)).
fof(t82_enumset1, axiom, ![X1]:k2_enumset1(X1,X1,X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/42a5356dbbf6454093949020b48842d0.E_input', t82_enumset1)).
fof(t61_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k4_enumset1(X1,X2,X3,X4,X5,X6),k1_tarski(X7)), file('data/MPTP/proofs/42a5356dbbf6454093949020b48842d0.E_input', t61_enumset1)).
fof(t50_enumset1, axiom, ![X1, X2, X3, X4, X5]:k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_tarski(X5)), file('data/MPTP/proofs/42a5356dbbf6454093949020b48842d0.E_input', t50_enumset1)).
fof(t66_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_enumset1(X6,X7,X8)), file('data/MPTP/proofs/42a5356dbbf6454093949020b48842d0.E_input', t66_enumset1)).
fof(t92_enumset1, conjecture, ![X1, X2]:k5_enumset1(X1,X1,X1,X1,X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/42a5356dbbf6454093949020b48842d0.E_input', t92_enumset1)).
fof(c_0_12, plain, ![X45, X46, X47, X48, X49, X50, X51]:k5_enumset1(X45,X46,X47,X48,X49,X50,X51)=k2_xboole_0(k1_enumset1(X45,X46,X47),k2_enumset1(X48,X49,X50,X51)), inference(variable_rename,[status(thm)],[t58_enumset1])).
fof(c_0_13, plain, ![X26, X27, X28]:k2_enumset1(X26,X26,X27,X28)=k1_enumset1(X26,X27,X28), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_14, plain, ![X37, X38, X39, X40, X41, X42, X43, X44]:k6_enumset1(X37,X38,X39,X40,X41,X42,X43,X44)=k2_xboole_0(k4_enumset1(X37,X38,X39,X40,X41,X42),k2_tarski(X43,X44)), inference(variable_rename,[status(thm)],[t67_enumset1])).
fof(c_0_15, plain, ![X35, X36]:k3_enumset1(X35,X35,X35,X35,X36)=k2_tarski(X35,X36), inference(variable_rename,[status(thm)],[t83_enumset1])).
fof(c_0_16, plain, ![X29, X30, X31, X32, X33]:k4_enumset1(X29,X29,X30,X31,X32,X33)=k3_enumset1(X29,X30,X31,X32,X33), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_17, plain, ![X52, X53, X54, X55, X56, X57, X58, X59]:k6_enumset1(X52,X53,X54,X55,X56,X57,X58,X59)=k2_xboole_0(k2_enumset1(X52,X53,X54,X55),k2_enumset1(X56,X57,X58,X59)), inference(variable_rename,[status(thm)],[l75_enumset1])).
fof(c_0_18, plain, ![X34]:k4_enumset1(X34,X34,X34,X34,X34,X34)=k1_tarski(X34), inference(variable_rename,[status(thm)],[t91_enumset1])).
fof(c_0_19, plain, ![X65]:k2_enumset1(X65,X65,X65,X65)=k1_tarski(X65), inference(variable_rename,[status(thm)],[t82_enumset1])).
fof(c_0_20, plain, ![X19, X20, X21, X22, X23, X24, X25]:k5_enumset1(X19,X20,X21,X22,X23,X24,X25)=k2_xboole_0(k4_enumset1(X19,X20,X21,X22,X23,X24),k1_tarski(X25)), inference(variable_rename,[status(thm)],[t61_enumset1])).
cnf(c_0_21, plain, (k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k1_enumset1(X1,X2,X3),k2_enumset1(X4,X5,X6,X7))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k4_enumset1(X1,X2,X3,X4,X5,X6),k2_tarski(X7,X8))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_24, plain, (k3_enumset1(X1,X1,X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27, plain, (k4_enumset1(X1,X1,X1,X1,X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, plain, (k2_enumset1(X1,X1,X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_29, plain, ![X60, X61, X62, X63, X64]:k3_enumset1(X60,X61,X62,X63,X64)=k2_xboole_0(k2_enumset1(X60,X61,X62,X63),k1_tarski(X64)), inference(variable_rename,[status(thm)],[t50_enumset1])).
cnf(c_0_30, plain, (k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k4_enumset1(X1,X2,X3,X4,X5,X6),k1_tarski(X7))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_31, plain, (k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k2_enumset1(X1,X1,X2,X3),k2_enumset1(X4,X5,X6,X7))), inference(rw,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_32, plain, (k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8))=k2_xboole_0(k4_enumset1(X1,X2,X3,X4,X5,X6),k4_enumset1(X7,X7,X7,X7,X7,X8))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_24]), c_0_25]), c_0_26])).
cnf(c_0_33, plain, (k4_enumset1(X1,X1,X1,X1,X1,X1)=k2_enumset1(X1,X1,X1,X1)), inference(rw,[status(thm)],[c_0_27, c_0_28])).
fof(c_0_34, plain, ![X11, X12, X13, X14, X15, X16, X17, X18]:k6_enumset1(X11,X12,X13,X14,X15,X16,X17,X18)=k2_xboole_0(k3_enumset1(X11,X12,X13,X14,X15),k1_enumset1(X16,X17,X18)), inference(variable_rename,[status(thm)],[t66_enumset1])).
cnf(c_0_35, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_tarski(X5))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, plain, (k2_xboole_0(k2_enumset1(X1,X1,X2,X3),k2_enumset1(X4,X5,X6,X7))=k2_xboole_0(k4_enumset1(X1,X2,X3,X4,X5,X6),k2_enumset1(X7,X7,X7,X7))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_28]), c_0_31])).
cnf(c_0_37, plain, (k2_xboole_0(k4_enumset1(X1,X2,X3,X4,X5,X6),k2_enumset1(X7,X7,X7,X7))=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X7))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
fof(c_0_38, negated_conjecture, ~(![X1, X2]:k5_enumset1(X1,X1,X1,X1,X1,X1,X2)=k2_tarski(X1,X2)), inference(assume_negation,[status(cth)],[t92_enumset1])).
cnf(c_0_39, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_enumset1(X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_40, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X5,X5,X5))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_28]), c_0_25])).
cnf(c_0_41, plain, (k2_xboole_0(k2_enumset1(X1,X1,X2,X3),k2_enumset1(X4,X5,X6,X7))=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X7))), inference(rw,[status(thm)],[c_0_36, c_0_37])).
fof(c_0_42, negated_conjecture, k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)!=k2_tarski(esk1_0,esk2_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_38])])])).
cnf(c_0_43, plain, (k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8))=k2_xboole_0(k4_enumset1(X1,X1,X2,X3,X4,X5),k2_enumset1(X6,X6,X7,X8))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_22]), c_0_25]), c_0_26])).
cnf(c_0_44, plain, (k2_xboole_0(k2_enumset1(X1,X1,X2,X3),k2_enumset1(X4,X5,X5,X5))=k4_enumset1(X1,X1,X2,X3,X4,X5)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_45, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)!=k2_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_46, plain, (k2_xboole_0(k2_enumset1(X1,X1,X1,X1),k2_enumset1(X2,X2,X3,X4))=k2_xboole_0(k2_enumset1(X1,X1,X1,X1),k2_enumset1(X1,X2,X3,X4))), inference(spm,[status(thm)],[c_0_43, c_0_33])).
cnf(c_0_47, plain, (k2_xboole_0(k2_enumset1(X1,X1,X1,X2),k2_enumset1(X3,X4,X5,X5))=k4_enumset1(X1,X1,X2,X3,X4,X5)), inference(spm,[status(thm)],[c_0_41, c_0_44])).
cnf(c_0_48, negated_conjecture, (k2_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),k2_enumset1(esk1_0,esk1_0,esk1_0,esk2_0))!=k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_24]), c_0_25]), c_0_31])).
cnf(c_0_49, plain, (k2_xboole_0(k2_enumset1(X1,X1,X1,X1),k2_enumset1(X1,X2,X2,X3))=k4_enumset1(X1,X1,X1,X1,X2,X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_41]), c_0_47])).
cnf(c_0_50, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_49])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 51
# Proof object clause steps            : 26
# Proof object formula steps           : 25
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 19
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 12
# Removed in clause preprocessing      : 6
# Initial clauses in saturation        : 6
# Processed clauses                    : 116
# ...of these trivial                  : 1
# ...subsumed                          : 78
# ...remaining for further processing  : 37
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 3
# Generated clauses                    : 610
# ...of the previous two non-trivial   : 523
# Contextual simplify-reflections      : 0
# Paramodulations                      : 610
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 27
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 18
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 406
# ...number of literals in the above   : 406
# Current number of archived formulas  : 0
# Current number of archived clauses   : 16
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 109
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 481
# BW rewrite match successes           : 167
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4098

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.003 s
# Total time               : 0.011 s
# Maximum resident set size: 28676 pages
