# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/42c8cb3a26674de68618187f7fe8ad54.E_input', d4_xboole_0)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/42c8cb3a26674de68618187f7fe8ad54.E_input', t3_xboole_0)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/42c8cb3a26674de68618187f7fe8ad54.E_input', d5_xboole_0)).
fof(t89_xboole_1, conjecture, ![X1, X2]:r1_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2)), file('data/MPTP/proofs/42c8cb3a26674de68618187f7fe8ad54.E_input', t89_xboole_1)).
fof(c_0_4, plain, ![X30, X31, X32, X33, X34, X35, X36, X37]:((((r2_hidden(X33,X30)|~r2_hidden(X33,X32)|X32!=k3_xboole_0(X30,X31))&(r2_hidden(X33,X31)|~r2_hidden(X33,X32)|X32!=k3_xboole_0(X30,X31)))&(~r2_hidden(X34,X30)|~r2_hidden(X34,X31)|r2_hidden(X34,X32)|X32!=k3_xboole_0(X30,X31)))&((~r2_hidden(esk6_3(X35,X36,X37),X37)|(~r2_hidden(esk6_3(X35,X36,X37),X35)|~r2_hidden(esk6_3(X35,X36,X37),X36))|X37=k3_xboole_0(X35,X36))&((r2_hidden(esk6_3(X35,X36,X37),X35)|r2_hidden(esk6_3(X35,X36,X37),X37)|X37=k3_xboole_0(X35,X36))&(r2_hidden(esk6_3(X35,X36,X37),X36)|r2_hidden(esk6_3(X35,X36,X37),X37)|X37=k3_xboole_0(X35,X36))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_5, plain, ![X39, X40, X42, X43, X44]:(((r2_hidden(esk7_2(X39,X40),X39)|r1_xboole_0(X39,X40))&(r2_hidden(esk7_2(X39,X40),X40)|r1_xboole_0(X39,X40)))&(~r2_hidden(X44,X42)|~r2_hidden(X44,X43)|~r1_xboole_0(X42,X43))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_6, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_7, plain, ![X14, X15, X16, X17, X18, X19, X20, X21]:((((r2_hidden(X17,X14)|~r2_hidden(X17,X16)|X16!=k4_xboole_0(X14,X15))&(~r2_hidden(X17,X15)|~r2_hidden(X17,X16)|X16!=k4_xboole_0(X14,X15)))&(~r2_hidden(X18,X14)|r2_hidden(X18,X15)|r2_hidden(X18,X16)|X16!=k4_xboole_0(X14,X15)))&((~r2_hidden(esk3_3(X19,X20,X21),X21)|(~r2_hidden(esk3_3(X19,X20,X21),X19)|r2_hidden(esk3_3(X19,X20,X21),X20))|X21=k4_xboole_0(X19,X20))&((r2_hidden(esk3_3(X19,X20,X21),X19)|r2_hidden(esk3_3(X19,X20,X21),X21)|X21=k4_xboole_0(X19,X20))&(~r2_hidden(esk3_3(X19,X20,X21),X20)|r2_hidden(esk3_3(X19,X20,X21),X21)|X21=k4_xboole_0(X19,X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_8, negated_conjecture, ~(![X1, X2]:r1_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))), inference(assume_negation,[status(cth)],[t89_xboole_1])).
cnf(c_0_9, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10, plain, (r2_hidden(esk7_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_6])).
cnf(c_0_12, plain, (r2_hidden(esk7_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_13, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_14, negated_conjecture, ~r1_xboole_0(k3_xboole_0(esk1_0,esk2_0),k4_xboole_0(esk1_0,esk2_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
cnf(c_0_15, plain, (r1_xboole_0(X1,X2)|~r2_hidden(esk7_2(X1,X2),X3)|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_16, plain, (r2_hidden(esk7_2(k3_xboole_0(X1,X2),X3),X2)|r1_xboole_0(k3_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_17, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_13])).
cnf(c_0_18, negated_conjecture, (~r1_xboole_0(k3_xboole_0(esk1_0,esk2_0),k4_xboole_0(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (r1_xboole_0(k3_xboole_0(X1,X2),X3)|~r1_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_20, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r2_hidden(esk7_2(X1,k4_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_17, c_0_10])).
cnf(c_0_21, negated_conjecture, (~r1_xboole_0(esk2_0,k4_xboole_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_22, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_20, c_0_12])).
cnf(c_0_23, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_22])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 24
# Proof object clause steps            : 15
# Proof object formula steps           : 9
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 6
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 4
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 31
# Processed clauses                    : 228
# ...of these trivial                  : 5
# ...subsumed                          : 123
# ...remaining for further processing  : 100
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 550
# ...of the previous two non-trivial   : 443
# Contextual simplify-reflections      : 1
# Paramodulations                      : 538
# Factorizations                       : 6
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 92
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 70
# Current number of unprocessed clauses: 245
# ...number of literals in the above   : 627
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2
# Clause-clause subsumption calls (NU) : 1478
# Rec. Clause-clause subsumption calls : 1165
# Non-unit clause-clause subsumptions  : 106
# Unit Clause-clause subsumption calls : 31
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 19
# BW rewrite match successes           : 8
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6037

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28716 pages
