# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.007 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t17_finset_1, conjecture, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(v1_finset_1(X1)=>v1_finset_1(k9_relat_1(X2,X1)))), file('data/MPTP/proofs/42cb37a897d74b689bf35a592fcfb2cf.E_input', t17_finset_1)).
fof(fc13_finset_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finset_1(X2))=>v1_finset_1(k9_relat_1(X1,X2))), file('data/MPTP/proofs/42cb37a897d74b689bf35a592fcfb2cf.E_input', fc13_finset_1)).
fof(c_0_2, negated_conjecture, ~(![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(v1_finset_1(X1)=>v1_finset_1(k9_relat_1(X2,X1))))), inference(assume_negation,[status(cth)],[t17_finset_1])).
fof(c_0_3, negated_conjecture, ((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&(v1_finset_1(esk1_0)&~v1_finset_1(k9_relat_1(esk2_0,esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
fof(c_0_4, plain, ![X5, X6]:(~v1_relat_1(X5)|~v1_funct_1(X5)|~v1_finset_1(X6)|v1_finset_1(k9_relat_1(X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc13_finset_1])])).
cnf(c_0_5, negated_conjecture, (~v1_finset_1(k9_relat_1(esk2_0,esk1_0))), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_6, plain, (v1_finset_1(k9_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, negated_conjecture, (v1_finset_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_8, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_9, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_10, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5, c_0_6]), c_0_7]), c_0_8]), c_0_9])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 11
# Proof object clause steps            : 6
# Proof object formula steps           : 5
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 5
# Proof object initial formulas used   : 2
# Proof object generating inferences   : 1
# Proof object simplifying inferences  : 4
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 2
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 5
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 5
# Processed clauses                    : 10
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 10
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 1
# ...of the previous two non-trivial   : 0
# Contextual simplify-reflections      : 0
# Paramodulations                      : 1
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 5
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 1
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 5
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 249

# -------------------------------------------------
# User time                : 0.005 s
# System time              : 0.002 s
# Total time               : 0.007 s
# Maximum resident set size: 28652 pages
