# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/43b2315bfca44541a9667278d39272bd.E_input', d5_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/43b2315bfca44541a9667278d39272bd.E_input', commutativity_k3_xboole_0)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/43b2315bfca44541a9667278d39272bd.E_input', t48_xboole_1)).
fof(t53_pre_topc, conjecture, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v3_pre_topc(X2,X1)=>k2_pre_topc(X1,k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))&((v2_pre_topc(X1)&k2_pre_topc(X1,k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=>v3_pre_topc(X2,X1))))), file('data/MPTP/proofs/43b2315bfca44541a9667278d39272bd.E_input', t53_pre_topc)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/43b2315bfca44541a9667278d39272bd.E_input', d3_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/43b2315bfca44541a9667278d39272bd.E_input', dt_l1_pre_topc)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/43b2315bfca44541a9667278d39272bd.E_input', t3_subset)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/43b2315bfca44541a9667278d39272bd.E_input', t4_subset_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/43b2315bfca44541a9667278d39272bd.E_input', t28_xboole_1)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/43b2315bfca44541a9667278d39272bd.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/43b2315bfca44541a9667278d39272bd.E_input', d4_subset_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/43b2315bfca44541a9667278d39272bd.E_input', t3_boole)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/43b2315bfca44541a9667278d39272bd.E_input', d10_xboole_0)).
fof(d6_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v4_pre_topc(X2,X1)<=>v3_pre_topc(k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2),X1)))), file('data/MPTP/proofs/43b2315bfca44541a9667278d39272bd.E_input', d6_pre_topc)).
fof(redefinition_k7_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k7_subset_1(X1,X2,X3)=k4_xboole_0(X2,X3)), file('data/MPTP/proofs/43b2315bfca44541a9667278d39272bd.E_input', redefinition_k7_subset_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/43b2315bfca44541a9667278d39272bd.E_input', t36_xboole_1)).
fof(t52_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v4_pre_topc(X2,X1)=>k2_pre_topc(X1,X2)=X2)&((v2_pre_topc(X1)&k2_pre_topc(X1,X2)=X2)=>v4_pre_topc(X2,X1))))), file('data/MPTP/proofs/43b2315bfca44541a9667278d39272bd.E_input', t52_pre_topc)).
fof(c_0_17, plain, ![X5, X6, X7, X8, X9, X10, X11, X12]:((((r2_hidden(X8,X5)|~r2_hidden(X8,X7)|X7!=k4_xboole_0(X5,X6))&(~r2_hidden(X8,X6)|~r2_hidden(X8,X7)|X7!=k4_xboole_0(X5,X6)))&(~r2_hidden(X9,X5)|r2_hidden(X9,X6)|r2_hidden(X9,X7)|X7!=k4_xboole_0(X5,X6)))&((~r2_hidden(esk1_3(X10,X11,X12),X12)|(~r2_hidden(esk1_3(X10,X11,X12),X10)|r2_hidden(esk1_3(X10,X11,X12),X11))|X12=k4_xboole_0(X10,X11))&((r2_hidden(esk1_3(X10,X11,X12),X10)|r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k4_xboole_0(X10,X11))&(~r2_hidden(esk1_3(X10,X11,X12),X11)|r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k4_xboole_0(X10,X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_18, plain, ![X14, X15]:k3_xboole_0(X14,X15)=k3_xboole_0(X15,X14), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_19, plain, ![X16, X17]:k4_xboole_0(X16,k4_xboole_0(X16,X17))=k3_xboole_0(X16,X17), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_20, negated_conjecture, ~(![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v3_pre_topc(X2,X1)=>k2_pre_topc(X1,k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))&((v2_pre_topc(X1)&k2_pre_topc(X1,k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=>v3_pre_topc(X2,X1)))))), inference(assume_negation,[status(cth)],[t53_pre_topc])).
fof(c_0_21, plain, ![X20]:(~l1_struct_0(X20)|k2_struct_0(X20)=u1_struct_0(X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
fof(c_0_22, plain, ![X21]:(~l1_pre_topc(X21)|l1_struct_0(X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_23, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_26, plain, ![X22, X23]:((~m1_subset_1(X22,k1_zfmisc_1(X23))|r1_tarski(X22,X23))&(~r1_tarski(X22,X23)|m1_subset_1(X22,k1_zfmisc_1(X23)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_27, plain, ![X24]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X24)), inference(variable_rename,[status(thm)],[t4_subset_1])).
fof(c_0_28, plain, ![X25, X26]:(~r1_tarski(X25,X26)|k3_xboole_0(X25,X26)=X25), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_29, negated_conjecture, (l1_pre_topc(esk2_0)&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&((((v2_pre_topc(esk2_0)|v3_pre_topc(esk3_0,esk2_0))&(k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0))=k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0)|v3_pre_topc(esk3_0,esk2_0)))&(~v3_pre_topc(esk3_0,esk2_0)|v3_pre_topc(esk3_0,esk2_0)))&(((v2_pre_topc(esk2_0)|k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0))!=k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0))&(k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0))=k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0)|k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0))!=k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0)))&(~v3_pre_topc(esk3_0,esk2_0)|k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0))!=k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])])).
cnf(c_0_30, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_32, plain, ![X27]:m1_subset_1(k2_subset_1(X27),k1_zfmisc_1(X27)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_33, plain, ![X28]:k2_subset_1(X28)=X28, inference(variable_rename,[status(thm)],[d4_subset_1])).
cnf(c_0_34, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_23])).
cnf(c_0_35, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25]), c_0_25])).
fof(c_0_36, plain, ![X29]:k4_xboole_0(X29,k1_xboole_0)=X29, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_37, plain, ![X30, X31]:(((r1_tarski(X30,X31)|X30!=X31)&(r1_tarski(X31,X30)|X30!=X31))&(~r1_tarski(X30,X31)|~r1_tarski(X31,X30)|X30=X31)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_38, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_39, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_40, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_41, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_42, plain, ![X32, X33]:((~v4_pre_topc(X33,X32)|v3_pre_topc(k7_subset_1(u1_struct_0(X32),k2_struct_0(X32),X33),X32)|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32)))|~l1_pre_topc(X32))&(~v3_pre_topc(k7_subset_1(u1_struct_0(X32),k2_struct_0(X32),X33),X32)|v4_pre_topc(X33,X32)|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32)))|~l1_pre_topc(X32))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_pre_topc])])])])).
cnf(c_0_43, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_44, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_45, plain, ![X34, X35, X36]:(~m1_subset_1(X35,k1_zfmisc_1(X34))|k7_subset_1(X34,X35,X36)=k4_xboole_0(X35,X36)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_subset_1])])).
cnf(c_0_46, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_47, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_48, plain, (~r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))|~r2_hidden(X1,k4_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_49, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_50, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_51, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
fof(c_0_52, plain, ![X37, X38]:r1_tarski(k4_xboole_0(X37,X38),X37), inference(variable_rename,[status(thm)],[t36_xboole_1])).
cnf(c_0_53, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=X1|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_40, c_0_25])).
cnf(c_0_54, negated_conjecture, (r1_tarski(esk3_0,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_38, c_0_41])).
cnf(c_0_55, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_56, plain, (v4_pre_topc(X2,X1)|~v3_pre_topc(k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_57, negated_conjecture, (k2_struct_0(esk2_0)=u1_struct_0(esk2_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_58, plain, (k7_subset_1(X2,X1,X3)=k4_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_59, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_60, plain, (~r2_hidden(X1,k4_xboole_0(k1_xboole_0,X2))|~r2_hidden(X1,k4_xboole_0(X2,X2))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_61, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_62, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_63, negated_conjecture, (k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,u1_struct_0(esk2_0)))=esk3_0), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_64, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_55])).
cnf(c_0_65, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_66, negated_conjecture, (v4_pre_topc(X1,esk2_0)|~v3_pre_topc(k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_44])])).
cnf(c_0_67, plain, (k7_subset_1(X1,X1,X2)=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_68, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_49]), c_0_49])])).
cnf(c_0_69, plain, (r2_hidden(esk1_3(X1,X2,X3),X1)|r2_hidden(esk1_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_70, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_71, negated_conjecture, (~r2_hidden(X1,k4_xboole_0(esk3_0,u1_struct_0(esk2_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_63]), c_0_64])).
cnf(c_0_72, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_65])).
cnf(c_0_73, negated_conjecture, (v4_pre_topc(X1,esk2_0)|~v3_pre_topc(k4_xboole_0(u1_struct_0(esk2_0),X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(rw,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_74, plain, (X1=k1_xboole_0|r2_hidden(esk1_3(k1_xboole_0,X2,X1),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70])).
cnf(c_0_75, plain, (r2_hidden(esk1_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk1_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_76, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_77, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|r2_hidden(esk1_3(X1,X2,k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
fof(c_0_78, plain, ![X39, X40]:((~v4_pre_topc(X40,X39)|k2_pre_topc(X39,X40)=X40|~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X39)))|~l1_pre_topc(X39))&(~v2_pre_topc(X39)|k2_pre_topc(X39,X40)!=X40|v4_pre_topc(X40,X39)|~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X39)))|~l1_pre_topc(X39))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t52_pre_topc])])])])).
cnf(c_0_79, negated_conjecture, (v4_pre_topc(k4_xboole_0(u1_struct_0(esk2_0),X1),esk2_0)|~v3_pre_topc(k4_xboole_0(X1,k4_xboole_0(X1,u1_struct_0(esk2_0))),esk2_0)|~m1_subset_1(k4_xboole_0(u1_struct_0(esk2_0),X1),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_73, c_0_35])).
cnf(c_0_80, negated_conjecture, (k4_xboole_0(esk3_0,u1_struct_0(esk2_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_71, c_0_74])).
cnf(c_0_81, negated_conjecture, (k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0))=k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0)|v3_pre_topc(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_82, plain, (v3_pre_topc(k7_subset_1(u1_struct_0(X2),k2_struct_0(X2),X1),X2)|~v4_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_83, negated_conjecture, (X1=k4_xboole_0(X2,u1_struct_0(esk2_0))|r2_hidden(esk1_3(X2,u1_struct_0(esk2_0),X1),X1)|~r2_hidden(esk1_3(X2,u1_struct_0(esk2_0),X1),esk3_0)), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_84, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k1_xboole_0|r2_hidden(esk1_3(k4_xboole_0(X1,X2),X3,k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_64, c_0_77])).
cnf(c_0_85, negated_conjecture, (~v3_pre_topc(esk3_0,esk2_0)|k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0))!=k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_86, plain, (k2_pre_topc(X2,X1)=X1|~v4_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_87, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_88, negated_conjecture, (v4_pre_topc(k4_xboole_0(u1_struct_0(esk2_0),esk3_0),esk2_0)|~v3_pre_topc(esk3_0,esk2_0)|~m1_subset_1(k4_xboole_0(u1_struct_0(esk2_0),esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_49])).
cnf(c_0_89, negated_conjecture, (k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),esk3_0))=k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),esk3_0)|v3_pre_topc(esk3_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_57]), c_0_57])).
cnf(c_0_90, negated_conjecture, (v3_pre_topc(k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),X1),esk2_0)|~v4_pre_topc(X1,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_57]), c_0_44])])).
cnf(c_0_91, negated_conjecture, (k4_xboole_0(k4_xboole_0(esk3_0,X1),u1_struct_0(esk2_0))=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_68])).
cnf(c_0_92, negated_conjecture, (k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),esk3_0))!=k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),esk3_0)|~v3_pre_topc(esk3_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_57]), c_0_57])).
cnf(c_0_93, plain, (k2_pre_topc(X1,X2)=X2|~v4_pre_topc(X2,X1)|~r1_tarski(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_94, negated_conjecture, (v4_pre_topc(k4_xboole_0(u1_struct_0(esk2_0),esk3_0),esk2_0)|~v3_pre_topc(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_87]), c_0_62])])).
cnf(c_0_95, negated_conjecture, (k2_pre_topc(esk2_0,k4_xboole_0(u1_struct_0(esk2_0),esk3_0))=k4_xboole_0(u1_struct_0(esk2_0),esk3_0)|v3_pre_topc(esk3_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_67]), c_0_67])).
cnf(c_0_96, negated_conjecture, (v2_pre_topc(esk2_0)|k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0))!=k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_97, negated_conjecture, (v3_pre_topc(k4_xboole_0(u1_struct_0(esk2_0),X1),esk2_0)|~v4_pre_topc(X1,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(rw,[status(thm)],[c_0_90, c_0_67])).
cnf(c_0_98, negated_conjecture, (k4_xboole_0(u1_struct_0(esk2_0),k4_xboole_0(u1_struct_0(esk2_0),k4_xboole_0(esk3_0,X1)))=k4_xboole_0(esk3_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_91]), c_0_49])).
cnf(c_0_99, negated_conjecture, (k2_pre_topc(esk2_0,k4_xboole_0(u1_struct_0(esk2_0),esk3_0))!=k4_xboole_0(u1_struct_0(esk2_0),esk3_0)|~v3_pre_topc(esk3_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_67]), c_0_67])).
cnf(c_0_100, negated_conjecture, (k2_pre_topc(esk2_0,k4_xboole_0(u1_struct_0(esk2_0),esk3_0))=k4_xboole_0(u1_struct_0(esk2_0),esk3_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_62]), c_0_44])]), c_0_95])).
cnf(c_0_101, negated_conjecture, (v2_pre_topc(esk2_0)|k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),esk3_0))!=k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_96, c_0_57]), c_0_57])).
cnf(c_0_102, negated_conjecture, (v3_pre_topc(k4_xboole_0(esk3_0,X1),esk2_0)|~v4_pre_topc(k4_xboole_0(u1_struct_0(esk2_0),k4_xboole_0(esk3_0,X1)),esk2_0)|~m1_subset_1(k4_xboole_0(u1_struct_0(esk2_0),k4_xboole_0(esk3_0,X1)),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_103, negated_conjecture, (~v3_pre_topc(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_99, c_0_100])])).
cnf(c_0_104, plain, (v4_pre_topc(X2,X1)|~v2_pre_topc(X1)|k2_pre_topc(X1,X2)!=X2|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_105, negated_conjecture, (v2_pre_topc(esk2_0)|k2_pre_topc(esk2_0,k4_xboole_0(u1_struct_0(esk2_0),esk3_0))!=k4_xboole_0(u1_struct_0(esk2_0),esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_101, c_0_67]), c_0_67])).
cnf(c_0_106, negated_conjecture, (~v4_pre_topc(k4_xboole_0(u1_struct_0(esk2_0),esk3_0),esk2_0)|~m1_subset_1(k4_xboole_0(u1_struct_0(esk2_0),esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_49]), c_0_103])).
cnf(c_0_107, plain, (v4_pre_topc(X1,X2)|k2_pre_topc(X2,X1)!=X1|~r1_tarski(X1,u1_struct_0(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_104, c_0_87])).
cnf(c_0_108, negated_conjecture, (v2_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_105, c_0_100])])).
cnf(c_0_109, negated_conjecture, (~v4_pre_topc(k4_xboole_0(u1_struct_0(esk2_0),esk3_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_87]), c_0_62])])).
cnf(c_0_110, negated_conjecture, (v4_pre_topc(X1,esk2_0)|k2_pre_topc(esk2_0,X1)!=X1|~r1_tarski(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_44])])).
cnf(c_0_111, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_100]), c_0_62])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 112
# Proof object clause steps            : 77
# Proof object formula steps           : 35
# Proof object conjectures             : 38
# Proof object clause conjectures      : 35
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 33
# Proof object simplifying inferences  : 50
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 30
# Processed clauses                    : 361
# ...of these trivial                  : 14
# ...subsumed                          : 168
# ...remaining for further processing  : 179
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 8
# Backward-rewritten                   : 20
# Generated clauses                    : 931
# ...of the previous two non-trivial   : 739
# Contextual simplify-reflections      : 3
# Paramodulations                      : 919
# Factorizations                       : 6
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 116
#    Positive orientable unit clauses  : 27
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 85
# Current number of unprocessed clauses: 413
# ...number of literals in the above   : 1100
# Current number of archived formulas  : 0
# Current number of archived clauses   : 60
# Clause-clause subsumption calls (NU) : 1985
# Rec. Clause-clause subsumption calls : 1706
# Non-unit clause-clause subsumptions  : 111
# Unit Clause-clause subsumption calls : 102
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 36
# BW rewrite match successes           : 13
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 15421

# -------------------------------------------------
# User time                : 0.019 s
# System time              : 0.000 s
# Total time               : 0.019 s
# Maximum resident set size: 28624 pages
