# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t152_zfmisc_1, axiom, ![X1, X2]:~(r2_hidden(X1,k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/43bc32c2c4f844b0a62ec53890be5ed8.E_input', t152_zfmisc_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/43bc32c2c4f844b0a62ec53890be5ed8.E_input', t113_zfmisc_1)).
fof(t7_tarski, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&r2_hidden(X4,X3))))))), file('data/MPTP/proofs/43bc32c2c4f844b0a62ec53890be5ed8.E_input', t7_tarski)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/43bc32c2c4f844b0a62ec53890be5ed8.E_input', t3_xboole_0)).
fof(t1_mcart_1, conjecture, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&r1_xboole_0(X2,X1))))), file('data/MPTP/proofs/43bc32c2c4f844b0a62ec53890be5ed8.E_input', t1_mcart_1)).
fof(d5_setfam_1, axiom, ![X1, X2, X3]:(X3=k3_setfam_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k3_xboole_0(X5,X6))))), file('data/MPTP/proofs/43bc32c2c4f844b0a62ec53890be5ed8.E_input', d5_setfam_1)).
fof(c_0_6, plain, ![X7, X8]:~r2_hidden(X7,k2_zfmisc_1(X7,X8)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_zfmisc_1])])).
fof(c_0_7, plain, ![X9, X10]:((k2_zfmisc_1(X9,X10)!=k1_xboole_0|(X9=k1_xboole_0|X10=k1_xboole_0))&((X9!=k1_xboole_0|k2_zfmisc_1(X9,X10)=k1_xboole_0)&(X10!=k1_xboole_0|k2_zfmisc_1(X9,X10)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_8, plain, (~epred1_0<=>![X2]:X2!=k1_xboole_0), introduced(definition)).
fof(c_0_9, plain, ![X11, X12, X14]:((r2_hidden(esk1_2(X11,X12),X12)|~r2_hidden(X11,X12))&(~r2_hidden(X14,X12)|~r2_hidden(X14,esk1_2(X11,X12))|~r2_hidden(X11,X12))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tarski])])])])])).
fof(c_0_10, plain, ![X32, X33, X35, X36, X37]:(((r2_hidden(esk7_2(X32,X33),X32)|r1_xboole_0(X32,X33))&(r2_hidden(esk7_2(X32,X33),X33)|r1_xboole_0(X32,X33)))&(~r2_hidden(X37,X35)|~r2_hidden(X37,X36)|~r1_xboole_0(X35,X36))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_11, plain, (~epred2_0<=>![X1]:~r2_hidden(X1,k1_xboole_0)), introduced(definition)).
cnf(c_0_12, plain, (~r2_hidden(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14, plain, (epred1_0|X1!=k1_xboole_0), inference(split_equiv,[status(thm)],[c_0_8])).
fof(c_0_15, negated_conjecture, ~(![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&r1_xboole_0(X2,X1)))))), inference(assume_negation,[status(cth)],[t1_mcart_1])).
cnf(c_0_16, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,esk1_2(X3,X2))|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, plain, (r2_hidden(esk7_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, plain, (~epred2_0|~epred1_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_8]), c_0_11])).
cnf(c_0_19, plain, (epred1_0), inference(er,[status(thm)],[c_0_14])).
fof(c_0_20, negated_conjecture, ![X39]:(esk8_0!=k1_xboole_0&(~r2_hidden(X39,esk8_0)|~r1_xboole_0(X39,esk8_0))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_15])])])])])).
cnf(c_0_21, plain, (r1_xboole_0(esk1_2(X1,X2),X3)|~r2_hidden(esk7_2(esk1_2(X1,X2),X3),X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_22, plain, (r2_hidden(esk7_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_23, plain, (epred2_0|~r2_hidden(X1,k1_xboole_0)), inference(split_equiv,[status(thm)],[c_0_11])).
cnf(c_0_24, plain, (r2_hidden(esk1_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_25, plain, (~epred2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_19])])).
fof(c_0_26, plain, ![X15, X16, X17, X18, X21, X22, X23, X24, X25, X26, X28, X29]:((((r2_hidden(esk2_4(X15,X16,X17,X18),X15)|~r2_hidden(X18,X17)|X17!=k3_setfam_1(X15,X16))&((r2_hidden(esk3_4(X15,X16,X17,X18),X16)|~r2_hidden(X18,X17)|X17!=k3_setfam_1(X15,X16))&(X18=k3_xboole_0(esk2_4(X15,X16,X17,X18),esk3_4(X15,X16,X17,X18))|~r2_hidden(X18,X17)|X17!=k3_setfam_1(X15,X16))))&(~r2_hidden(X22,X15)|(~r2_hidden(X23,X16)|X21!=k3_xboole_0(X22,X23))|r2_hidden(X21,X17)|X17!=k3_setfam_1(X15,X16)))&((~r2_hidden(esk4_3(X24,X25,X26),X26)|(~r2_hidden(X28,X24)|(~r2_hidden(X29,X25)|esk4_3(X24,X25,X26)!=k3_xboole_0(X28,X29)))|X26=k3_setfam_1(X24,X25))&((r2_hidden(esk5_3(X24,X25,X26),X24)|r2_hidden(esk4_3(X24,X25,X26),X26)|X26=k3_setfam_1(X24,X25))&((r2_hidden(esk6_3(X24,X25,X26),X25)|r2_hidden(esk4_3(X24,X25,X26),X26)|X26=k3_setfam_1(X24,X25))&(esk4_3(X24,X25,X26)=k3_xboole_0(esk5_3(X24,X25,X26),esk6_3(X24,X25,X26))|r2_hidden(esk4_3(X24,X25,X26),X26)|X26=k3_setfam_1(X24,X25)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_setfam_1])])])])])])).
cnf(c_0_27, negated_conjecture, (~r2_hidden(X1,esk8_0)|~r1_xboole_0(X1,esk8_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (r1_xboole_0(esk1_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_29, plain, (~r2_hidden(X1,k1_xboole_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
cnf(c_0_30, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r2_hidden(esk4_3(X1,X2,X3),X3)|X3=k3_setfam_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_31, negated_conjecture, (~r2_hidden(X1,esk8_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_24])).
cnf(c_0_32, plain, (X1=k3_setfam_1(X2,k1_xboole_0)|r2_hidden(esk4_3(X2,k1_xboole_0,X1),X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_33, negated_conjecture, (k3_setfam_1(X1,k1_xboole_0)=esk8_0), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_34, plain, (X1=esk8_0|r2_hidden(esk4_3(X2,k1_xboole_0,X1),X1)), inference(rw,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_35, negated_conjecture, (esk8_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_36, plain, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_34]), c_0_25]), c_0_35]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 37
# Proof object clause steps            : 22
# Proof object formula steps           : 15
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 9
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 19
# Processed clauses                    : 81
# ...of these trivial                  : 2
# ...subsumed                          : 9
# ...remaining for further processing  : 70
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 9
# Generated clauses                    : 113
# ...of the previous two non-trivial   : 110
# Contextual simplify-reflections      : 1
# Paramodulations                      : 106
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 41
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 8
#    Non-unit-clauses                  : 25
# Current number of unprocessed clauses: 64
# ...number of literals in the above   : 159
# Current number of archived formulas  : 0
# Current number of archived clauses   : 28
# Clause-clause subsumption calls (NU) : 120
# Rec. Clause-clause subsumption calls : 77
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 47
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2353

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 26172 pages
