# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_yellow_0(X1,X2)=>(X3=k2_yellow_0(X1,X2)<=>(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))))))), file('data/MPTP/proofs/4447f9fda5854b9da290f070fbd3a73c.E_input', d10_yellow_0)).
fof(t31_yellow_0, conjecture, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3))=>(r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2)))))&((r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2))))=>(X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3)))))), file('data/MPTP/proofs/4447f9fda5854b9da290f070fbd3a73c.E_input', t31_yellow_0)).
fof(t16_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3))))))), file('data/MPTP/proofs/4447f9fda5854b9da290f070fbd3a73c.E_input', t16_yellow_0)).
fof(c_0_3, plain, ![X17, X18, X19, X20]:(((r1_lattice3(X17,X18,X19)|X19!=k2_yellow_0(X17,X18)|~r2_yellow_0(X17,X18)|~m1_subset_1(X19,u1_struct_0(X17))|~l1_orders_2(X17))&(~m1_subset_1(X20,u1_struct_0(X17))|(~r1_lattice3(X17,X18,X20)|r1_orders_2(X17,X20,X19))|X19!=k2_yellow_0(X17,X18)|~r2_yellow_0(X17,X18)|~m1_subset_1(X19,u1_struct_0(X17))|~l1_orders_2(X17)))&((m1_subset_1(esk7_3(X17,X18,X19),u1_struct_0(X17))|~r1_lattice3(X17,X18,X19)|X19=k2_yellow_0(X17,X18)|~r2_yellow_0(X17,X18)|~m1_subset_1(X19,u1_struct_0(X17))|~l1_orders_2(X17))&((r1_lattice3(X17,X18,esk7_3(X17,X18,X19))|~r1_lattice3(X17,X18,X19)|X19=k2_yellow_0(X17,X18)|~r2_yellow_0(X17,X18)|~m1_subset_1(X19,u1_struct_0(X17))|~l1_orders_2(X17))&(~r1_orders_2(X17,esk7_3(X17,X18,X19),X19)|~r1_lattice3(X17,X18,X19)|X19=k2_yellow_0(X17,X18)|~r2_yellow_0(X17,X18)|~m1_subset_1(X19,u1_struct_0(X17))|~l1_orders_2(X17))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_yellow_0])])])])])).
fof(c_0_4, negated_conjecture, ~(![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3))=>(r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2)))))&((r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2))))=>(X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3))))))), inference(assume_negation,[status(cth)],[t31_yellow_0])).
cnf(c_0_5, plain, (r1_lattice3(X1,X2,X3)|X3!=k2_yellow_0(X1,X2)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
fof(c_0_6, negated_conjecture, ![X9]:((v5_orders_2(esk1_0)&l1_orders_2(esk1_0))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(((((r1_lattice3(esk1_0,esk3_0,esk2_0)|esk2_0=k2_yellow_0(esk1_0,esk3_0))&(~m1_subset_1(X9,u1_struct_0(esk1_0))|(~r1_lattice3(esk1_0,esk3_0,X9)|r1_orders_2(esk1_0,X9,esk2_0))|esk2_0=k2_yellow_0(esk1_0,esk3_0)))&(esk2_0!=k2_yellow_0(esk1_0,esk3_0)|~r2_yellow_0(esk1_0,esk3_0)|esk2_0=k2_yellow_0(esk1_0,esk3_0)))&(((r1_lattice3(esk1_0,esk3_0,esk2_0)|r2_yellow_0(esk1_0,esk3_0))&(~m1_subset_1(X9,u1_struct_0(esk1_0))|(~r1_lattice3(esk1_0,esk3_0,X9)|r1_orders_2(esk1_0,X9,esk2_0))|r2_yellow_0(esk1_0,esk3_0)))&(esk2_0!=k2_yellow_0(esk1_0,esk3_0)|~r2_yellow_0(esk1_0,esk3_0)|r2_yellow_0(esk1_0,esk3_0))))&((((r1_lattice3(esk1_0,esk3_0,esk2_0)|(m1_subset_1(esk4_0,u1_struct_0(esk1_0))|~r1_lattice3(esk1_0,esk3_0,esk2_0)))&(~m1_subset_1(X9,u1_struct_0(esk1_0))|(~r1_lattice3(esk1_0,esk3_0,X9)|r1_orders_2(esk1_0,X9,esk2_0))|(m1_subset_1(esk4_0,u1_struct_0(esk1_0))|~r1_lattice3(esk1_0,esk3_0,esk2_0))))&(esk2_0!=k2_yellow_0(esk1_0,esk3_0)|~r2_yellow_0(esk1_0,esk3_0)|(m1_subset_1(esk4_0,u1_struct_0(esk1_0))|~r1_lattice3(esk1_0,esk3_0,esk2_0))))&((((r1_lattice3(esk1_0,esk3_0,esk2_0)|(r1_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk1_0,esk3_0,esk2_0)))&(~m1_subset_1(X9,u1_struct_0(esk1_0))|(~r1_lattice3(esk1_0,esk3_0,X9)|r1_orders_2(esk1_0,X9,esk2_0))|(r1_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk1_0,esk3_0,esk2_0))))&(esk2_0!=k2_yellow_0(esk1_0,esk3_0)|~r2_yellow_0(esk1_0,esk3_0)|(r1_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk1_0,esk3_0,esk2_0))))&(((r1_lattice3(esk1_0,esk3_0,esk2_0)|(~r1_orders_2(esk1_0,esk4_0,esk2_0)|~r1_lattice3(esk1_0,esk3_0,esk2_0)))&(~m1_subset_1(X9,u1_struct_0(esk1_0))|(~r1_lattice3(esk1_0,esk3_0,X9)|r1_orders_2(esk1_0,X9,esk2_0))|(~r1_orders_2(esk1_0,esk4_0,esk2_0)|~r1_lattice3(esk1_0,esk3_0,esk2_0))))&(esk2_0!=k2_yellow_0(esk1_0,esk3_0)|~r2_yellow_0(esk1_0,esk3_0)|(~r1_orders_2(esk1_0,esk4_0,esk2_0)|~r1_lattice3(esk1_0,esk3_0,esk2_0))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).
fof(c_0_7, plain, ![X10, X11, X13, X14, X15]:(((m1_subset_1(esk5_2(X10,X11),u1_struct_0(X10))|~r2_yellow_0(X10,X11)|(~v5_orders_2(X10)|~l1_orders_2(X10)))&((r1_lattice3(X10,X11,esk5_2(X10,X11))|~r2_yellow_0(X10,X11)|(~v5_orders_2(X10)|~l1_orders_2(X10)))&(~m1_subset_1(X13,u1_struct_0(X10))|(~r1_lattice3(X10,X11,X13)|r1_orders_2(X10,X13,esk5_2(X10,X11)))|~r2_yellow_0(X10,X11)|(~v5_orders_2(X10)|~l1_orders_2(X10)))))&((m1_subset_1(esk6_3(X10,X14,X15),u1_struct_0(X10))|~r1_lattice3(X10,X14,X15)|~m1_subset_1(X15,u1_struct_0(X10))|r2_yellow_0(X10,X14)|(~v5_orders_2(X10)|~l1_orders_2(X10)))&((r1_lattice3(X10,X14,esk6_3(X10,X14,X15))|~r1_lattice3(X10,X14,X15)|~m1_subset_1(X15,u1_struct_0(X10))|r2_yellow_0(X10,X14)|(~v5_orders_2(X10)|~l1_orders_2(X10)))&(~r1_orders_2(X10,esk6_3(X10,X14,X15),X15)|~r1_lattice3(X10,X14,X15)|~m1_subset_1(X15,u1_struct_0(X10))|r2_yellow_0(X10,X14)|(~v5_orders_2(X10)|~l1_orders_2(X10)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_yellow_0])])])])])])).
cnf(c_0_8, plain, (r1_lattice3(X1,X2,k2_yellow_0(X1,X2))|~r2_yellow_0(X1,X2)|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_5])).
cnf(c_0_9, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk2_0)|esk2_0=k2_yellow_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk2_0)|r2_yellow_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13, plain, (r2_yellow_0(X1,X2)|~r1_orders_2(X1,esk6_3(X1,X2,X3),X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|r2_yellow_0(esk1_0,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_lattice3(esk1_0,esk3_0,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_16, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_lattice3(esk1_0,esk3_0,X1)|~r1_lattice3(esk1_0,esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_17, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10]), c_0_11])]), c_0_12])).
cnf(c_0_18, negated_conjecture, (r2_yellow_0(esk1_0,esk3_0)|r2_yellow_0(esk1_0,X1)|~r1_lattice3(esk1_0,esk3_0,esk6_3(esk1_0,X1,esk2_0))|~r1_lattice3(esk1_0,X1,esk2_0)|~m1_subset_1(esk6_3(esk1_0,X1,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_10]), c_0_11]), c_0_15])])).
cnf(c_0_19, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_20, plain, (X3=k2_yellow_0(X1,X2)|~r1_orders_2(X1,esk7_3(X1,X2,X3),X3)|~r1_lattice3(X1,X2,X3)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_21, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_17])])).
cnf(c_0_22, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk4_0)|esk2_0!=k2_yellow_0(esk1_0,esk3_0)|~r2_yellow_0(esk1_0,esk3_0)|~r1_lattice3(esk1_0,esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_23, negated_conjecture, (r2_yellow_0(esk1_0,esk3_0)|r2_yellow_0(esk1_0,X1)|~r1_lattice3(esk1_0,esk3_0,esk6_3(esk1_0,X1,esk2_0))|~r1_lattice3(esk1_0,X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_10]), c_0_11]), c_0_15])])).
cnf(c_0_24, plain, (r1_lattice3(X1,X2,esk6_3(X1,X2,X3))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_25, negated_conjecture, (k2_yellow_0(esk1_0,X1)=esk2_0|r1_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk1_0,esk3_0,esk7_3(esk1_0,X1,esk2_0))|~r1_lattice3(esk1_0,X1,esk2_0)|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(esk7_3(esk1_0,X1,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_10]), c_0_11])])).
cnf(c_0_26, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|X3=k2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_27, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk4_0)|k2_yellow_0(esk1_0,esk3_0)!=esk2_0|~r2_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_17])])).
cnf(c_0_28, negated_conjecture, (r2_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_17]), c_0_10]), c_0_11]), c_0_15])])).
cnf(c_0_29, negated_conjecture, (k2_yellow_0(esk1_0,X1)=esk2_0|r1_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk1_0,esk3_0,esk7_3(esk1_0,X1,esk2_0))|~r1_lattice3(esk1_0,X1,esk2_0)|~r2_yellow_0(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_10]), c_0_11])])).
cnf(c_0_30, plain, (r1_lattice3(X1,X2,esk7_3(X1,X2,X3))|X3=k2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_31, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk4_0)|k2_yellow_0(esk1_0,esk3_0)!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_28])])).
cnf(c_0_32, plain, (r1_orders_2(X2,X1,X4)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|X4!=k2_yellow_0(X2,X3)|~r2_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_33, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|esk2_0=k2_yellow_0(esk1_0,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_lattice3(esk1_0,esk3_0,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_34, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_17]), c_0_28]), c_0_10]), c_0_11])]), c_0_31])).
cnf(c_0_35, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|m1_subset_1(esk4_0,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_lattice3(esk1_0,esk3_0,X1)|~r1_lattice3(esk1_0,esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_36, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_lattice3(esk1_0,esk3_0,X1)|~r1_orders_2(esk1_0,esk4_0,esk2_0)|~r1_lattice3(esk1_0,esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_37, plain, (r1_orders_2(X1,X2,k2_yellow_0(X1,X3))|~r1_lattice3(X1,X3,X2)|~r2_yellow_0(X1,X3)|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_38, negated_conjecture, (k2_yellow_0(esk1_0,esk3_0)=esk2_0|r1_orders_2(esk1_0,esk4_0,esk2_0)|~m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_39, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|m1_subset_1(esk4_0,u1_struct_0(esk1_0))|~r1_lattice3(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_17])])).
cnf(c_0_40, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0)|~r1_lattice3(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_17])])).
cnf(c_0_41, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|~r1_lattice3(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_28]), c_0_10]), c_0_11])]), c_0_39]), c_0_40])).
cnf(c_0_42, negated_conjecture, (esk2_0!=k2_yellow_0(esk1_0,esk3_0)|~r2_yellow_0(esk1_0,esk3_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0)|~r1_lattice3(esk1_0,esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_43, negated_conjecture, (k2_yellow_0(esk1_0,X1)=esk2_0|~r1_lattice3(esk1_0,esk3_0,esk7_3(esk1_0,X1,esk2_0))|~r1_lattice3(esk1_0,X1,esk2_0)|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(esk7_3(esk1_0,X1,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_41]), c_0_10]), c_0_11])])).
cnf(c_0_44, negated_conjecture, (k2_yellow_0(esk1_0,esk3_0)!=esk2_0|~r1_orders_2(esk1_0,esk4_0,esk2_0)|~r2_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_17])])).
cnf(c_0_45, negated_conjecture, (k2_yellow_0(esk1_0,X1)=esk2_0|~r1_lattice3(esk1_0,esk3_0,esk7_3(esk1_0,X1,esk2_0))|~r1_lattice3(esk1_0,X1,esk2_0)|~r2_yellow_0(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_26]), c_0_10]), c_0_11])])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))|esk2_0!=k2_yellow_0(esk1_0,esk3_0)|~r2_yellow_0(esk1_0,esk3_0)|~r1_lattice3(esk1_0,esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_47, negated_conjecture, (k2_yellow_0(esk1_0,esk3_0)!=esk2_0|~r1_orders_2(esk1_0,esk4_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_28])])).
cnf(c_0_48, negated_conjecture, (k2_yellow_0(esk1_0,esk3_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_30]), c_0_17]), c_0_28]), c_0_10]), c_0_11])])).
cnf(c_0_49, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))|k2_yellow_0(esk1_0,esk3_0)!=esk2_0|~r2_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_17])])).
cnf(c_0_50, negated_conjecture, (~r1_orders_2(esk1_0,esk4_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_48])])).
cnf(c_0_51, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))|k2_yellow_0(esk1_0,esk3_0)!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_28])])).
cnf(c_0_52, negated_conjecture, (~m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_41]), c_0_34])])).
cnf(c_0_53, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_48])]), c_0_52]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 54
# Proof object clause steps            : 47
# Proof object formula steps           : 7
# Proof object conjectures             : 40
# Proof object clause conjectures      : 37
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 73
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 24
# Processed clauses                    : 92
# ...of these trivial                  : 1
# ...subsumed                          : 0
# ...remaining for further processing  : 90
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 14
# Backward-rewritten                   : 23
# Generated clauses                    : 65
# ...of the previous two non-trivial   : 60
# Contextual simplify-reflections      : 9
# Paramodulations                      : 63
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 27
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 17
# Current number of unprocessed clauses: 9
# ...number of literals in the above   : 55
# Current number of archived formulas  : 0
# Current number of archived clauses   : 61
# Clause-clause subsumption calls (NU) : 321
# Rec. Clause-clause subsumption calls : 36
# Non-unit clause-clause subsumptions  : 20
# Unit Clause-clause subsumption calls : 10
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 11
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4364

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28560 pages
