# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t2_yellow19, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(X1)))&(v2_waybel_0(X2,k3_yellow_1(X1))&(v13_waybel_0(X2,k3_yellow_1(X1))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X1))))))))=>![X3]:~((r2_hidden(X3,X2)&v1_xboole_0(X3))))), file('data/MPTP/proofs/44bd85ca2b5c48a8a7a71fe5165a686a.E_input', t2_yellow19)).
fof(t4_yellow_1, axiom, ![X1]:k3_yellow_1(X1)=k2_yellow_1(k9_setfam_1(X1)), file('data/MPTP/proofs/44bd85ca2b5c48a8a7a71fe5165a686a.E_input', t4_yellow_1)).
fof(redefinition_k9_setfam_1, axiom, ![X1]:k9_setfam_1(X1)=k1_zfmisc_1(X1), file('data/MPTP/proofs/44bd85ca2b5c48a8a7a71fe5165a686a.E_input', redefinition_k9_setfam_1)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/44bd85ca2b5c48a8a7a71fe5165a686a.E_input', d1_zfmisc_1)).
fof(t11_waybel_7, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X1))))=>(v13_waybel_0(X2,k3_yellow_1(X1))<=>![X3, X4]:((r1_tarski(X3,X4)&(r1_tarski(X4,X1)&r2_hidden(X3,X2)))=>r2_hidden(X4,X2)))), file('data/MPTP/proofs/44bd85ca2b5c48a8a7a71fe5165a686a.E_input', t11_waybel_7)).
fof(t1_yellow_1, axiom, ![X1]:(u1_struct_0(k2_yellow_1(X1))=X1&u1_orders_2(k2_yellow_1(X1))=k1_yellow_1(X1)), file('data/MPTP/proofs/44bd85ca2b5c48a8a7a71fe5165a686a.E_input', t1_yellow_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/44bd85ca2b5c48a8a7a71fe5165a686a.E_input', t2_xboole_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/44bd85ca2b5c48a8a7a71fe5165a686a.E_input', d1_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/44bd85ca2b5c48a8a7a71fe5165a686a.E_input', d3_tarski)).
fof(t49_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(![X3]:(m1_subset_1(X3,X1)=>r2_hidden(X3,X2))=>X1=X2)), file('data/MPTP/proofs/44bd85ca2b5c48a8a7a71fe5165a686a.E_input', t49_subset_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/44bd85ca2b5c48a8a7a71fe5165a686a.E_input', t2_subset)).
fof(d7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(v1_subset_1(X2,X1)<=>~(X2=X1))), file('data/MPTP/proofs/44bd85ca2b5c48a8a7a71fe5165a686a.E_input', d7_subset_1)).
fof(rc3_subset_1, axiom, ![X1]:?[X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))&~(v1_subset_1(X2,X1))), file('data/MPTP/proofs/44bd85ca2b5c48a8a7a71fe5165a686a.E_input', rc3_subset_1)).
fof(c_0_13, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(X1)))&(v2_waybel_0(X2,k3_yellow_1(X1))&(v13_waybel_0(X2,k3_yellow_1(X1))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X1))))))))=>![X3]:~((r2_hidden(X3,X2)&v1_xboole_0(X3)))))), inference(assume_negation,[status(cth)],[t2_yellow19])).
fof(c_0_14, plain, ![X24]:k3_yellow_1(X24)=k2_yellow_1(k9_setfam_1(X24)), inference(variable_rename,[status(thm)],[t4_yellow_1])).
fof(c_0_15, plain, ![X23]:k9_setfam_1(X23)=k1_zfmisc_1(X23), inference(variable_rename,[status(thm)],[redefinition_k9_setfam_1])).
fof(c_0_16, negated_conjecture, (~v1_xboole_0(esk1_0)&((~v1_xboole_0(esk2_0)&(v1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(esk1_0)))&(v2_waybel_0(esk2_0,k3_yellow_1(esk1_0))&(v13_waybel_0(esk2_0,k3_yellow_1(esk1_0))&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(esk1_0))))))))&(r2_hidden(esk3_0,esk2_0)&v1_xboole_0(esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_13])])])])).
cnf(c_0_17, plain, (k3_yellow_1(X1)=k2_yellow_1(k9_setfam_1(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (k9_setfam_1(X1)=k1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_19, plain, ![X16, X17, X18, X19, X20, X21]:(((~r2_hidden(X18,X17)|r1_tarski(X18,X16)|X17!=k1_zfmisc_1(X16))&(~r1_tarski(X19,X16)|r2_hidden(X19,X17)|X17!=k1_zfmisc_1(X16)))&((~r2_hidden(esk5_2(X20,X21),X21)|~r1_tarski(esk5_2(X20,X21),X20)|X21=k1_zfmisc_1(X20))&(r2_hidden(esk5_2(X20,X21),X21)|r1_tarski(esk5_2(X20,X21),X20)|X21=k1_zfmisc_1(X20)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
fof(c_0_20, plain, ![X36, X37, X38, X39]:((~v13_waybel_0(X37,k3_yellow_1(X36))|(~r1_tarski(X38,X39)|(~r1_tarski(X39,X36)|~r2_hidden(X38,X37))|r2_hidden(X39,X37))|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X36)))))&(((r1_tarski(esk8_2(X36,X37),esk9_2(X36,X37))|v13_waybel_0(X37,k3_yellow_1(X36))|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X36)))))&((r1_tarski(esk9_2(X36,X37),X36)|v13_waybel_0(X37,k3_yellow_1(X36))|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X36)))))&(r2_hidden(esk8_2(X36,X37),X37)|v13_waybel_0(X37,k3_yellow_1(X36))|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X36)))))))&(~r2_hidden(esk9_2(X36,X37),X37)|v13_waybel_0(X37,k3_yellow_1(X36))|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X36))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_waybel_7])])])])])).
cnf(c_0_21, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (k3_yellow_1(X1)=k2_yellow_1(k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_17, c_0_18])).
fof(c_0_23, plain, ![X11]:(u1_struct_0(k2_yellow_1(X11))=X11&u1_orders_2(k2_yellow_1(X11))=k1_yellow_1(X11)), inference(variable_rename,[status(thm)],[t1_yellow_1])).
cnf(c_0_24, plain, (r2_hidden(X1,X3)|~r1_tarski(X1,X2)|X3!=k1_zfmisc_1(X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_25, plain, ![X31]:r1_tarski(k1_xboole_0,X31), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_26, plain, (r2_hidden(X4,X1)|~v13_waybel_0(X1,k3_yellow_1(X2))|~r1_tarski(X3,X4)|~r1_tarski(X4,X2)|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_27, plain, ![X32, X33, X34]:((~v1_xboole_0(X32)|~r2_hidden(X33,X32))&(r2_hidden(esk7_1(X34),X34)|v1_xboole_0(X34))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_28, plain, ![X25, X26, X27, X28, X29]:((~r1_tarski(X25,X26)|(~r2_hidden(X27,X25)|r2_hidden(X27,X26)))&((r2_hidden(esk6_2(X28,X29),X28)|r1_tarski(X28,X29))&(~r2_hidden(esk6_2(X28,X29),X29)|r1_tarski(X28,X29)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_29, plain, ![X8, X9]:((m1_subset_1(esk4_2(X8,X9),X8)|X8=X9|~m1_subset_1(X9,k1_zfmisc_1(X8)))&(~r2_hidden(esk4_2(X8,X9),X9)|X8=X9|~m1_subset_1(X9,k1_zfmisc_1(X8)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t49_subset_1])])])])).
cnf(c_0_30, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k1_zfmisc_1(esk1_0)))))), inference(rw,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_31, plain, (u1_struct_0(k2_yellow_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, plain, (r2_hidden(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(er,[status(thm)],[c_0_24])).
cnf(c_0_33, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, plain, (r2_hidden(X4,X1)|~r2_hidden(X3,X1)|~r1_tarski(X4,X2)|~r1_tarski(X3,X4)|~v13_waybel_0(X1,k2_yellow_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k1_zfmisc_1(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_22]), c_0_22])).
cnf(c_0_35, negated_conjecture, (v13_waybel_0(esk2_0,k3_yellow_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_36, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_38, plain, ![X12, X13]:(~m1_subset_1(X12,X13)|(v1_xboole_0(X13)|r2_hidden(X12,X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_39, plain, (m1_subset_1(esk4_2(X1,X2),X1)|X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_40, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))), inference(rw,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_41, plain, (r2_hidden(k1_xboole_0,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_42, plain, (r2_hidden(X1,X2)|~r1_tarski(X1,X3)|~r1_tarski(X4,X1)|~r2_hidden(X4,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))|~v13_waybel_0(X2,k2_yellow_1(k1_zfmisc_1(X3)))), inference(rw,[status(thm)],[c_0_34, c_0_31])).
cnf(c_0_43, negated_conjecture, (v13_waybel_0(esk2_0,k2_yellow_1(k1_zfmisc_1(esk1_0)))), inference(rw,[status(thm)],[c_0_35, c_0_22])).
cnf(c_0_44, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_45, negated_conjecture, (v1_xboole_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_46, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_47, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_48, negated_conjecture, (k1_zfmisc_1(esk1_0)=esk2_0|m1_subset_1(esk4_2(k1_zfmisc_1(esk1_0),esk2_0),k1_zfmisc_1(esk1_0))), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_49, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_36, c_0_41])).
fof(c_0_50, plain, ![X14, X15]:((~v1_subset_1(X15,X14)|X15!=X14|~m1_subset_1(X15,k1_zfmisc_1(X14)))&(X15=X14|v1_subset_1(X15,X14)|~m1_subset_1(X15,k1_zfmisc_1(X14)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_subset_1])])])])).
fof(c_0_51, plain, ![X42]:(m1_subset_1(esk10_1(X42),k1_zfmisc_1(X42))&~v1_subset_1(esk10_1(X42),X42)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[rc3_subset_1])])])).
cnf(c_0_52, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_53, negated_conjecture, (r2_hidden(X1,esk2_0)|~r1_tarski(X1,esk1_0)|~r1_tarski(X2,X1)|~r2_hidden(X2,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_40])])).
cnf(c_0_54, negated_conjecture, (r1_tarski(esk3_0,X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_55, negated_conjecture, (r2_hidden(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_56, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(er,[status(thm)],[c_0_46])).
cnf(c_0_57, negated_conjecture, (k1_zfmisc_1(esk1_0)=esk2_0|r2_hidden(esk4_2(k1_zfmisc_1(esk1_0),esk2_0),k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])).
cnf(c_0_58, plain, (X1=X2|~r2_hidden(esk4_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_59, plain, (X1=X2|v1_subset_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_60, plain, (m1_subset_1(esk10_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_61, plain, (~v1_subset_1(esk10_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_62, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(k2_yellow_1(k1_zfmisc_1(esk1_0))))), inference(rw,[status(thm)],[c_0_52, c_0_22])).
cnf(c_0_63, negated_conjecture, (r2_hidden(X1,esk2_0)|~r1_tarski(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55])])).
cnf(c_0_64, negated_conjecture, (k1_zfmisc_1(esk1_0)=esk2_0|r1_tarski(esk4_2(k1_zfmisc_1(esk1_0),esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_65, negated_conjecture, (k1_zfmisc_1(esk1_0)=esk2_0|~r2_hidden(esk4_2(k1_zfmisc_1(esk1_0),esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_58, c_0_40])).
cnf(c_0_66, plain, (esk10_1(X1)=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61])).
cnf(c_0_67, negated_conjecture, (v1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))), inference(rw,[status(thm)],[c_0_62, c_0_31])).
cnf(c_0_68, negated_conjecture, (k1_zfmisc_1(esk1_0)=esk2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_65])).
cnf(c_0_69, plain, (~v1_subset_1(X1,X1)), inference(rw,[status(thm)],[c_0_61, c_0_66])).
cnf(c_0_70, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_68]), c_0_69]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 71
# Proof object clause steps            : 44
# Proof object formula steps           : 27
# Proof object conjectures             : 22
# Proof object clause conjectures      : 19
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 21
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 31
# Processed clauses                    : 156
# ...of these trivial                  : 0
# ...subsumed                          : 34
# ...remaining for further processing  : 122
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 25
# Generated clauses                    : 179
# ...of the previous two non-trivial   : 134
# Contextual simplify-reflections      : 1
# Paramodulations                      : 175
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 60
#    Positive orientable unit clauses  : 16
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 40
# Current number of unprocessed clauses: 34
# ...number of literals in the above   : 110
# Current number of archived formulas  : 0
# Current number of archived clauses   : 62
# Clause-clause subsumption calls (NU) : 758
# Rec. Clause-clause subsumption calls : 473
# Non-unit clause-clause subsumptions  : 19
# Unit Clause-clause subsumption calls : 145
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 15
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4133

# -------------------------------------------------
# User time                : 0.004 s
# System time              : 0.008 s
# Total time               : 0.012 s
# Maximum resident set size: 28572 pages
