# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', t2_xboole_1)).
fof(fraenkel_a_2_0_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&(v3_orders_2(X2)&(v4_orders_2(X2)&(v5_orders_2(X2)&l1_orders_2(X2)))))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r2_hidden(X1,a_2_0_orders_2(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r2_hidden(X5,X3)=>r2_orders_2(X2,X5,X4))))))), file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', fraenkel_a_2_0_orders_2)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', d4_subset_1)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', t4_subset_1)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', t4_subset)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', d3_tarski)).
fof(t43_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>k1_orders_2(X1,k1_struct_0(X1))=u1_struct_0(X1)), file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', t43_orders_2)).
fof(d2_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k1_struct_0(X1)=k1_xboole_0), file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', d2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', dt_l1_orders_2)).
fof(d12_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k1_orders_2(X1,X2)=a_2_0_orders_2(X1,X2))), file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', d12_orders_2)).
fof(t44_orders_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>k1_orders_2(X1,k2_struct_0(X1))=k1_xboole_0), file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', t44_orders_2)).
fof(dt_k2_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', dt_k2_struct_0)).
fof(reflexivity_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>r3_orders_2(X1,X2,X2)), file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', reflexivity_r3_orders_2)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', t3_subset)).
fof(t30_orders_2, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~((r1_orders_2(X1,X2,X3)&r2_orders_2(X1,X3,X2)))))), file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', t30_orders_2)).
fof(t9_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k4_tarski(X3,X4))))))), file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', t9_mcart_1)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', redefinition_r3_orders_2)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/4521e2bc91cd4b20ae3c07a9077cc083.E_input', d3_struct_0)).
fof(c_0_20, plain, ![X23, X24]:(~r2_hidden(X23,X24)|~r1_tarski(X24,X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_21, plain, ![X27]:r1_tarski(k1_xboole_0,X27), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_22, plain, ![X28, X29, X30, X32, X33]:(((m1_subset_1(esk4_3(X28,X29,X30),u1_struct_0(X29))|~r2_hidden(X28,a_2_0_orders_2(X29,X30))|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))))&((X28=esk4_3(X28,X29,X30)|~r2_hidden(X28,a_2_0_orders_2(X29,X30))|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))))&(~m1_subset_1(X32,u1_struct_0(X29))|(~r2_hidden(X32,X30)|r2_orders_2(X29,X32,esk4_3(X28,X29,X30)))|~r2_hidden(X28,a_2_0_orders_2(X29,X30))|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))))))&((m1_subset_1(esk5_4(X28,X29,X30,X33),u1_struct_0(X29))|X28!=X33|~m1_subset_1(X33,u1_struct_0(X29))|r2_hidden(X28,a_2_0_orders_2(X29,X30))|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))))&((r2_hidden(esk5_4(X28,X29,X30,X33),X30)|X28!=X33|~m1_subset_1(X33,u1_struct_0(X29))|r2_hidden(X28,a_2_0_orders_2(X29,X30))|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))))&(~r2_orders_2(X29,esk5_4(X28,X29,X30,X33),X33)|X28!=X33|~m1_subset_1(X33,u1_struct_0(X29))|r2_hidden(X28,a_2_0_orders_2(X29,X30))|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))|~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_0_orders_2])])])])])])).
fof(c_0_23, plain, ![X13]:m1_subset_1(k2_subset_1(X13),k1_zfmisc_1(X13)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_24, plain, ![X26]:k2_subset_1(X26)=X26, inference(variable_rename,[status(thm)],[d4_subset_1])).
cnf(c_0_25, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (r2_hidden(esk5_4(X1,X2,X3,X4),X3)|r2_hidden(X1,a_2_0_orders_2(X2,X3))|v2_struct_0(X2)|X1!=X4|~m1_subset_1(X4,u1_struct_0(X2))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_28, plain, ![X22]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X22)), inference(variable_rename,[status(thm)],[t4_subset_1])).
fof(c_0_29, plain, ![X45, X46, X47]:(~r2_hidden(X45,X46)|~m1_subset_1(X46,k1_zfmisc_1(X47))|m1_subset_1(X45,X47)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_30, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_32, plain, ![X7, X8, X9, X10, X11]:((~r1_tarski(X7,X8)|(~r2_hidden(X9,X7)|r2_hidden(X9,X8)))&((r2_hidden(esk2_2(X10,X11),X10)|r1_tarski(X10,X11))&(~r2_hidden(esk2_2(X10,X11),X11)|r1_tarski(X10,X11)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_33, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_34, plain, (r2_hidden(esk5_4(X1,X2,X3,X1),X3)|r2_hidden(X1,a_2_0_orders_2(X2,X3))|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_35, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_30, c_0_31])).
fof(c_0_38, plain, ![X14]:(v2_struct_0(X14)|(~v3_orders_2(X14)|(~v4_orders_2(X14)|(~v5_orders_2(X14)|~l1_orders_2(X14))))|k1_orders_2(X14,k1_struct_0(X14))=u1_struct_0(X14)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t43_orders_2])])])).
fof(c_0_39, plain, ![X15]:(~l1_struct_0(X15)|k1_struct_0(X15)=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_struct_0])])).
fof(c_0_40, plain, ![X25]:(~l1_orders_2(X25)|l1_struct_0(X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_41, plain, (r1_tarski(X1,X2)|~r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_42, plain, (r2_hidden(X1,a_2_0_orders_2(X2,k1_xboole_0))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])])).
cnf(c_0_43, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_44, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_45, plain, ![X16, X17]:(v2_struct_0(X16)|(~v3_orders_2(X16)|(~v4_orders_2(X16)|(~v5_orders_2(X16)|~l1_orders_2(X16))))|(~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|k1_orders_2(X16,X17)=a_2_0_orders_2(X16,X17))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d12_orders_2])])])])).
cnf(c_0_46, plain, (v2_struct_0(X1)|k1_orders_2(X1,k1_struct_0(X1))=u1_struct_0(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_47, plain, (k1_struct_0(X1)=k1_xboole_0|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_48, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_49, plain, (r1_tarski(X1,a_2_0_orders_2(X2,k1_xboole_0))|v2_struct_0(X2)|~m1_subset_1(esk2_2(X1,a_2_0_orders_2(X2,k1_xboole_0)),u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_50, plain, (m1_subset_1(esk2_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_51, plain, (v2_struct_0(X1)|k1_orders_2(X1,X2)=a_2_0_orders_2(X1,X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_52, plain, (k1_orders_2(X1,k1_xboole_0)=u1_struct_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])).
fof(c_0_53, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>k1_orders_2(X1,k2_struct_0(X1))=k1_xboole_0)), inference(assume_negation,[status(cth)],[t44_orders_2])).
cnf(c_0_54, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_55, plain, (r1_tarski(u1_struct_0(X1),a_2_0_orders_2(X1,k1_xboole_0))|v2_struct_0(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_56, plain, (a_2_0_orders_2(X1,k1_xboole_0)=u1_struct_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_35])])).
fof(c_0_57, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&k1_orders_2(esk1_0,k2_struct_0(esk1_0))!=k1_xboole_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_53])])])])).
fof(c_0_58, plain, ![X50]:(~l1_struct_0(X50)|m1_subset_1(k2_struct_0(X50),k1_zfmisc_1(u1_struct_0(X50)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_struct_0])])).
cnf(c_0_59, plain, (r2_hidden(X1,a_2_0_orders_2(X2,k1_xboole_0))|v2_struct_0(X2)|~r2_hidden(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_60, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(spm,[status(thm)],[c_0_42, c_0_56])).
cnf(c_0_61, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_62, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_63, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_64, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_65, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_66, plain, (m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_67, plain, (r1_tarski(X1,a_2_0_orders_2(X2,k1_xboole_0))|v2_struct_0(X2)|~r2_hidden(esk2_2(X1,a_2_0_orders_2(X2,k1_xboole_0)),u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(spm,[status(thm)],[c_0_41, c_0_59])).
cnf(c_0_68, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62]), c_0_63]), c_0_64])]), c_0_65])).
cnf(c_0_69, plain, (m1_subset_1(X1,u1_struct_0(X2))|~l1_struct_0(X2)|~r2_hidden(X1,k2_struct_0(X2))), inference(spm,[status(thm)],[c_0_36, c_0_66])).
cnf(c_0_70, negated_conjecture, (r1_tarski(X1,a_2_0_orders_2(esk1_0,k1_xboole_0))|~m1_subset_1(esk2_2(X1,a_2_0_orders_2(esk1_0,k1_xboole_0)),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_62]), c_0_63]), c_0_61]), c_0_64])]), c_0_65])).
cnf(c_0_71, plain, (r1_tarski(X1,u1_struct_0(X2))|v2_struct_0(X2)|~m1_subset_1(esk2_2(X1,u1_struct_0(X2)),u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(spm,[status(thm)],[c_0_49, c_0_56])).
cnf(c_0_72, plain, (m1_subset_1(esk2_2(k2_struct_0(X1),X2),u1_struct_0(X1))|r1_tarski(k2_struct_0(X1),X2)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_69, c_0_44])).
fof(c_0_73, plain, ![X36, X37, X38]:(v2_struct_0(X36)|(~v3_orders_2(X36)|~l1_orders_2(X36))|(~m1_subset_1(X37,u1_struct_0(X36))|~m1_subset_1(X38,u1_struct_0(X36)))|r3_orders_2(X36,X37,X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r3_orders_2])])])).
cnf(c_0_74, negated_conjecture, (r1_tarski(u1_struct_0(esk1_0),a_2_0_orders_2(esk1_0,k1_xboole_0))), inference(spm,[status(thm)],[c_0_70, c_0_50])).
cnf(c_0_75, plain, (r1_tarski(k2_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_48])).
cnf(c_0_76, plain, (v2_struct_0(X1)|r3_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_77, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_orders_2(X2,X3))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_78, negated_conjecture, (r2_hidden(X1,a_2_0_orders_2(esk1_0,k1_xboole_0))|~r2_hidden(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_54, c_0_74])).
cnf(c_0_79, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k2_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(spm,[status(thm)],[c_0_54, c_0_75])).
cnf(c_0_80, plain, (r3_orders_2(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X4,a_2_0_orders_2(X1,X3))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_81, negated_conjecture, (r1_tarski(X1,a_2_0_orders_2(esk1_0,k1_xboole_0))|~r2_hidden(esk2_2(X1,a_2_0_orders_2(esk1_0,k1_xboole_0)),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_41, c_0_78])).
cnf(c_0_82, plain, (r2_hidden(esk2_2(k2_struct_0(X1),X2),u1_struct_0(X1))|r1_tarski(k2_struct_0(X1),X2)|v2_struct_0(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_79, c_0_44])).
cnf(c_0_83, negated_conjecture, (r3_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_78]), c_0_35]), c_0_62]), c_0_63]), c_0_61]), c_0_64])]), c_0_65])).
fof(c_0_84, plain, ![X48, X49]:((~m1_subset_1(X48,k1_zfmisc_1(X49))|r1_tarski(X48,X49))&(~r1_tarski(X48,X49)|m1_subset_1(X48,k1_zfmisc_1(X49)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_85, negated_conjecture, (r1_tarski(k2_struct_0(esk1_0),a_2_0_orders_2(esk1_0,k1_xboole_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_62]), c_0_63]), c_0_61]), c_0_64])]), c_0_65])).
cnf(c_0_86, plain, (X1=esk4_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_orders_2(X2,X3))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_87, plain, ![X39, X40, X41]:(~v5_orders_2(X39)|~l1_orders_2(X39)|(~m1_subset_1(X40,u1_struct_0(X39))|(~m1_subset_1(X41,u1_struct_0(X39))|(~r1_orders_2(X39,X40,X41)|~r2_orders_2(X39,X41,X40))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_orders_2])])])).
cnf(c_0_88, plain, (r2_orders_2(X2,X1,esk4_3(X4,X2,X3))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r2_hidden(X4,a_2_0_orders_2(X2,X3))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_89, negated_conjecture, (r3_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_83, c_0_68])).
fof(c_0_90, plain, ![X18, X20, X21]:((r2_hidden(esk3_1(X18),X18)|X18=k1_xboole_0)&((~r2_hidden(X20,X18)|esk3_1(X18)!=k4_tarski(X20,X21)|X18=k1_xboole_0)&(~r2_hidden(X21,X18)|esk3_1(X18)!=k4_tarski(X20,X21)|X18=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_mcart_1])])])])])])).
cnf(c_0_91, negated_conjecture, (k1_orders_2(esk1_0,k2_struct_0(esk1_0))!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_92, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_93, negated_conjecture, (r1_tarski(k2_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_56]), c_0_62]), c_0_63]), c_0_61]), c_0_64])]), c_0_65])).
cnf(c_0_94, plain, (esk4_3(esk2_2(a_2_0_orders_2(X1,X2),X3),X1,X2)=esk2_2(a_2_0_orders_2(X1,X2),X3)|r1_tarski(a_2_0_orders_2(X1,X2),X3)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_86, c_0_44])).
cnf(c_0_95, plain, (r2_hidden(X2,a_2_0_orders_2(X1,X3))|v2_struct_0(X1)|~r2_orders_2(X1,esk5_4(X2,X1,X3,X4),X4)|X2!=X4|~m1_subset_1(X4,u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_96, plain, (~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)|~r2_orders_2(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_97, plain, (r2_orders_2(X1,X2,esk4_3(X3,X1,X4))|v2_struct_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X3,a_2_0_orders_2(X1,X4))|~r2_hidden(X2,X4)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_88, c_0_36])).
fof(c_0_98, plain, ![X42, X43, X44]:((~r3_orders_2(X42,X43,X44)|r1_orders_2(X42,X43,X44)|(v2_struct_0(X42)|(~v3_orders_2(X42)|~l1_orders_2(X42))|(~m1_subset_1(X43,u1_struct_0(X42))|~m1_subset_1(X44,u1_struct_0(X42)))))&(~r1_orders_2(X42,X43,X44)|r3_orders_2(X42,X43,X44)|(v2_struct_0(X42)|(~v3_orders_2(X42)|~l1_orders_2(X42))|(~m1_subset_1(X43,u1_struct_0(X42))|~m1_subset_1(X44,u1_struct_0(X42)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_99, negated_conjecture, (r3_orders_2(esk1_0,X1,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X3,a_2_0_orders_2(esk1_0,X2))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_77]), c_0_62]), c_0_63]), c_0_61]), c_0_64])]), c_0_65])).
cnf(c_0_100, plain, (r2_hidden(esk3_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_101, negated_conjecture, (a_2_0_orders_2(esk1_0,k2_struct_0(esk1_0))!=k1_xboole_0|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_51]), c_0_62]), c_0_63]), c_0_61]), c_0_64])]), c_0_65])).
cnf(c_0_102, negated_conjecture, (m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_103, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~m1_subset_1(esk2_2(X1,u1_struct_0(esk1_0)),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_41, c_0_68])).
cnf(c_0_104, plain, (m1_subset_1(esk2_2(a_2_0_orders_2(X1,X2),X3),u1_struct_0(X1))|r1_tarski(a_2_0_orders_2(X1,X2),X3)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_94]), c_0_44])).
cnf(c_0_105, plain, (r2_hidden(X1,a_2_0_orders_2(X2,X3))|v2_struct_0(X2)|~r2_orders_2(X2,esk5_4(X1,X2,X3,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(er,[status(thm)],[c_0_95])).
cnf(c_0_106, plain, (v2_struct_0(X1)|~r1_orders_2(X1,esk4_3(X2,X1,X3),X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,a_2_0_orders_2(X1,X3))|~r2_hidden(X4,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_36]), c_0_77])).
cnf(c_0_107, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_108, negated_conjecture, (a_2_0_orders_2(esk1_0,X1)=k1_xboole_0|r3_orders_2(esk1_0,X2,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_109, negated_conjecture, (a_2_0_orders_2(esk1_0,k2_struct_0(esk1_0))!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_101, c_0_102])])).
cnf(c_0_110, negated_conjecture, (r1_tarski(a_2_0_orders_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_62]), c_0_63]), c_0_61]), c_0_64])]), c_0_65])).
cnf(c_0_111, plain, (r2_hidden(esk4_3(X1,X2,X3),a_2_0_orders_2(X2,X4))|v2_struct_0(X2)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(esk5_4(esk4_3(X1,X2,X3),X2,X4,esk4_3(X1,X2,X3)),X3)|~r2_hidden(X1,a_2_0_orders_2(X2,X3))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_97]), c_0_77])).
fof(c_0_112, plain, ![X35]:(~l1_struct_0(X35)|k2_struct_0(X35)=u1_struct_0(X35)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_113, plain, (v2_struct_0(X1)|~r3_orders_2(X1,esk4_3(X2,X1,X3),X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,a_2_0_orders_2(X1,X3))|~r2_hidden(X4,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_36]), c_0_77])).
cnf(c_0_114, negated_conjecture, (r3_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_102]), c_0_109])).
cnf(c_0_115, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X1,a_2_0_orders_2(esk1_0,X2))), inference(spm,[status(thm)],[c_0_54, c_0_110])).
cnf(c_0_116, plain, (r2_hidden(esk4_3(X1,X2,X3),a_2_0_orders_2(X2,X3))|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,a_2_0_orders_2(X2,X3))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_34]), c_0_77])).
cnf(c_0_117, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_118, negated_conjecture, (~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(esk4_3(X2,esk1_0,X1),X1)|~r2_hidden(X2,a_2_0_orders_2(esk1_0,X1))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_62]), c_0_63]), c_0_61]), c_0_64])]), c_0_65]), c_0_36])).
cnf(c_0_119, negated_conjecture, (r2_hidden(esk4_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X1,a_2_0_orders_2(esk1_0,X2))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_62]), c_0_63]), c_0_61]), c_0_64])]), c_0_65])).
cnf(c_0_120, negated_conjecture, (k1_orders_2(esk1_0,u1_struct_0(esk1_0))!=k1_xboole_0|~l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_91, c_0_117])).
cnf(c_0_121, negated_conjecture, (~r2_hidden(X1,a_2_0_orders_2(esk1_0,u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_37])])).
cnf(c_0_122, negated_conjecture, (a_2_0_orders_2(esk1_0,u1_struct_0(esk1_0))!=k1_xboole_0|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_51]), c_0_37]), c_0_62]), c_0_63]), c_0_61]), c_0_64])]), c_0_65])).
cnf(c_0_123, negated_conjecture, (a_2_0_orders_2(esk1_0,u1_struct_0(esk1_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_121, c_0_100])).
cnf(c_0_124, negated_conjecture, (~l1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_122, c_0_123])])).
cnf(c_0_125, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_48]), c_0_62])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 126
# Proof object clause steps            : 85
# Proof object formula steps           : 41
# Proof object conjectures             : 35
# Proof object clause conjectures      : 32
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 48
# Proof object simplifying inferences  : 94
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 20
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 36
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 35
# Processed clauses                    : 1278
# ...of these trivial                  : 0
# ...subsumed                          : 717
# ...remaining for further processing  : 561
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 75
# Backward-rewritten                   : 4
# Generated clauses                    : 2848
# ...of the previous two non-trivial   : 2707
# Contextual simplify-reflections      : 39
# Paramodulations                      : 2845
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 444
#    Positive orientable unit clauses  : 15
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 11
#    Non-unit-clauses                  : 418
# Current number of unprocessed clauses: 1399
# ...number of literals in the above   : 11093
# Current number of archived formulas  : 0
# Current number of archived clauses   : 115
# Clause-clause subsumption calls (NU) : 49317
# Rec. Clause-clause subsumption calls : 14424
# Non-unit clause-clause subsumptions  : 731
# Unit Clause-clause subsumption calls : 512
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 82105

# -------------------------------------------------
# User time                : 0.073 s
# System time              : 0.007 s
# Total time               : 0.080 s
# Maximum resident set size: 28724 pages
