# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_S0e
# and selection function SelectLargestNegLit.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t56_funct_1, conjecture, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((v2_funct_1(X2)&r2_hidden(X1,k1_relat_1(X2)))=>(X1=k1_funct_1(k2_funct_1(X2),k1_funct_1(X2,X1))&X1=k1_funct_1(k5_relat_1(X2,k2_funct_1(X2)),X1)))), file('data/MPTP/proofs/452c6da3181943fda05eca1acb798078.E_input', t56_funct_1)).
fof(dt_k2_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_relat_1(k2_funct_1(X1))&v1_funct_1(k2_funct_1(X1)))), file('data/MPTP/proofs/452c6da3181943fda05eca1acb798078.E_input', dt_k2_funct_1)).
fof(t55_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)=>(k2_relat_1(X1)=k1_relat_1(k2_funct_1(X1))&k1_relat_1(X1)=k2_relat_1(k2_funct_1(X1))))), file('data/MPTP/proofs/452c6da3181943fda05eca1acb798078.E_input', t55_funct_1)).
fof(t54_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(X2=k2_funct_1(X1)<=>(k1_relat_1(X2)=k2_relat_1(X1)&![X3, X4]:(((r2_hidden(X3,k2_relat_1(X1))&X4=k1_funct_1(X2,X3))=>(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4)))&((r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4))=>(r2_hidden(X3,k2_relat_1(X1))&X4=k1_funct_1(X2,X3))))))))), file('data/MPTP/proofs/452c6da3181943fda05eca1acb798078.E_input', t54_funct_1)).
fof(t182_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>k1_relat_1(k5_relat_1(X1,X2))=k10_relat_1(X1,k1_relat_1(X2)))), file('data/MPTP/proofs/452c6da3181943fda05eca1acb798078.E_input', t182_relat_1)).
fof(t169_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k10_relat_1(X1,k2_relat_1(X1))=k1_relat_1(X1)), file('data/MPTP/proofs/452c6da3181943fda05eca1acb798078.E_input', t169_relat_1)).
fof(t22_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k1_relat_1(k5_relat_1(X3,X2)))=>k1_funct_1(k5_relat_1(X3,X2),X1)=k1_funct_1(X2,k1_funct_1(X3,X1))))), file('data/MPTP/proofs/452c6da3181943fda05eca1acb798078.E_input', t22_funct_1)).
fof(c_0_7, negated_conjecture, ~(![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((v2_funct_1(X2)&r2_hidden(X1,k1_relat_1(X2)))=>(X1=k1_funct_1(k2_funct_1(X2),k1_funct_1(X2,X1))&X1=k1_funct_1(k5_relat_1(X2,k2_funct_1(X2)),X1))))), inference(assume_negation,[status(cth)],[t56_funct_1])).
fof(c_0_8, plain, ![X17]:((v1_relat_1(k2_funct_1(X17))|(~v1_relat_1(X17)|~v1_funct_1(X17)))&(v1_funct_1(k2_funct_1(X17))|(~v1_relat_1(X17)|~v1_funct_1(X17)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_funct_1])])])).
fof(c_0_9, negated_conjecture, ((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&((v2_funct_1(esk2_0)&r2_hidden(esk1_0,k1_relat_1(esk2_0)))&(esk1_0!=k1_funct_1(k2_funct_1(esk2_0),k1_funct_1(esk2_0,esk1_0))|esk1_0!=k1_funct_1(k5_relat_1(esk2_0,k2_funct_1(esk2_0)),esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
fof(c_0_10, plain, ![X24]:((k2_relat_1(X24)=k1_relat_1(k2_funct_1(X24))|~v2_funct_1(X24)|(~v1_relat_1(X24)|~v1_funct_1(X24)))&(k1_relat_1(X24)=k2_relat_1(k2_funct_1(X24))|~v2_funct_1(X24)|(~v1_relat_1(X24)|~v1_funct_1(X24)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_funct_1])])])).
fof(c_0_11, plain, ![X7, X8, X9, X10, X11, X12]:(((k1_relat_1(X8)=k2_relat_1(X7)|X8!=k2_funct_1(X7)|(~v1_relat_1(X8)|~v1_funct_1(X8))|~v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&(((r2_hidden(X10,k1_relat_1(X7))|(~r2_hidden(X9,k2_relat_1(X7))|X10!=k1_funct_1(X8,X9))|X8!=k2_funct_1(X7)|(~v1_relat_1(X8)|~v1_funct_1(X8))|~v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&(X9=k1_funct_1(X7,X10)|(~r2_hidden(X9,k2_relat_1(X7))|X10!=k1_funct_1(X8,X9))|X8!=k2_funct_1(X7)|(~v1_relat_1(X8)|~v1_funct_1(X8))|~v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7))))&((r2_hidden(X11,k2_relat_1(X7))|(~r2_hidden(X12,k1_relat_1(X7))|X11!=k1_funct_1(X7,X12))|X8!=k2_funct_1(X7)|(~v1_relat_1(X8)|~v1_funct_1(X8))|~v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&(X12=k1_funct_1(X8,X11)|(~r2_hidden(X12,k1_relat_1(X7))|X11!=k1_funct_1(X7,X12))|X8!=k2_funct_1(X7)|(~v1_relat_1(X8)|~v1_funct_1(X8))|~v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7))))))&(((((r2_hidden(esk6_2(X7,X8),k1_relat_1(X7))|r2_hidden(esk3_2(X7,X8),k2_relat_1(X7))|k1_relat_1(X8)!=k2_relat_1(X7)|X8=k2_funct_1(X7)|(~v1_relat_1(X8)|~v1_funct_1(X8))|~v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&(esk5_2(X7,X8)=k1_funct_1(X7,esk6_2(X7,X8))|r2_hidden(esk3_2(X7,X8),k2_relat_1(X7))|k1_relat_1(X8)!=k2_relat_1(X7)|X8=k2_funct_1(X7)|(~v1_relat_1(X8)|~v1_funct_1(X8))|~v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7))))&(~r2_hidden(esk5_2(X7,X8),k2_relat_1(X7))|esk6_2(X7,X8)!=k1_funct_1(X8,esk5_2(X7,X8))|r2_hidden(esk3_2(X7,X8),k2_relat_1(X7))|k1_relat_1(X8)!=k2_relat_1(X7)|X8=k2_funct_1(X7)|(~v1_relat_1(X8)|~v1_funct_1(X8))|~v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7))))&(((r2_hidden(esk6_2(X7,X8),k1_relat_1(X7))|esk4_2(X7,X8)=k1_funct_1(X8,esk3_2(X7,X8))|k1_relat_1(X8)!=k2_relat_1(X7)|X8=k2_funct_1(X7)|(~v1_relat_1(X8)|~v1_funct_1(X8))|~v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&(esk5_2(X7,X8)=k1_funct_1(X7,esk6_2(X7,X8))|esk4_2(X7,X8)=k1_funct_1(X8,esk3_2(X7,X8))|k1_relat_1(X8)!=k2_relat_1(X7)|X8=k2_funct_1(X7)|(~v1_relat_1(X8)|~v1_funct_1(X8))|~v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7))))&(~r2_hidden(esk5_2(X7,X8),k2_relat_1(X7))|esk6_2(X7,X8)!=k1_funct_1(X8,esk5_2(X7,X8))|esk4_2(X7,X8)=k1_funct_1(X8,esk3_2(X7,X8))|k1_relat_1(X8)!=k2_relat_1(X7)|X8=k2_funct_1(X7)|(~v1_relat_1(X8)|~v1_funct_1(X8))|~v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))))&(((r2_hidden(esk6_2(X7,X8),k1_relat_1(X7))|(~r2_hidden(esk4_2(X7,X8),k1_relat_1(X7))|esk3_2(X7,X8)!=k1_funct_1(X7,esk4_2(X7,X8)))|k1_relat_1(X8)!=k2_relat_1(X7)|X8=k2_funct_1(X7)|(~v1_relat_1(X8)|~v1_funct_1(X8))|~v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&(esk5_2(X7,X8)=k1_funct_1(X7,esk6_2(X7,X8))|(~r2_hidden(esk4_2(X7,X8),k1_relat_1(X7))|esk3_2(X7,X8)!=k1_funct_1(X7,esk4_2(X7,X8)))|k1_relat_1(X8)!=k2_relat_1(X7)|X8=k2_funct_1(X7)|(~v1_relat_1(X8)|~v1_funct_1(X8))|~v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7))))&(~r2_hidden(esk5_2(X7,X8),k2_relat_1(X7))|esk6_2(X7,X8)!=k1_funct_1(X8,esk5_2(X7,X8))|(~r2_hidden(esk4_2(X7,X8),k1_relat_1(X7))|esk3_2(X7,X8)!=k1_funct_1(X7,esk4_2(X7,X8)))|k1_relat_1(X8)!=k2_relat_1(X7)|X8=k2_funct_1(X7)|(~v1_relat_1(X8)|~v1_funct_1(X8))|~v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t54_funct_1])])])])])])).
fof(c_0_12, plain, ![X22, X23]:(~v1_relat_1(X22)|(~v1_relat_1(X23)|k1_relat_1(k5_relat_1(X22,X23))=k10_relat_1(X22,k1_relat_1(X23)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t182_relat_1])])])).
cnf(c_0_13, plain, (v1_relat_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, plain, (k2_relat_1(X1)=k1_relat_1(k2_funct_1(X1))|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, negated_conjecture, (v2_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_18, plain, ![X18]:(~v1_relat_1(X18)|k10_relat_1(X18,k2_relat_1(X18))=k1_relat_1(X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t169_relat_1])])).
cnf(c_0_19, plain, (X1=k1_funct_1(X2,X3)|~r2_hidden(X1,k1_relat_1(X4))|X3!=k1_funct_1(X4,X1)|X2!=k2_funct_1(X4)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X4)|~v1_relat_1(X4)|~v1_funct_1(X4)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_21, plain, (k1_relat_1(k5_relat_1(X1,X2))=k10_relat_1(X1,k1_relat_1(X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22, negated_conjecture, (v1_relat_1(k2_funct_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])])).
cnf(c_0_23, negated_conjecture, (k1_relat_1(k2_funct_1(esk2_0))=k2_relat_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_14]), c_0_15])])).
cnf(c_0_24, plain, (k10_relat_1(X1,k2_relat_1(X1))=k1_relat_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, negated_conjecture, (k1_funct_1(X1,X2)=esk1_0|X2!=k1_funct_1(esk2_0,esk1_0)|X1!=k2_funct_1(esk2_0)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_17]), c_0_14]), c_0_15])])).
cnf(c_0_26, plain, (v1_funct_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_27, plain, ![X19, X20, X21]:(~v1_relat_1(X20)|~v1_funct_1(X20)|(~v1_relat_1(X21)|~v1_funct_1(X21)|(~r2_hidden(X19,k1_relat_1(k5_relat_1(X21,X20)))|k1_funct_1(k5_relat_1(X21,X20),X19)=k1_funct_1(X20,k1_funct_1(X21,X19))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t22_funct_1])])])).
cnf(c_0_28, negated_conjecture, (k1_relat_1(k5_relat_1(X1,k2_funct_1(esk2_0)))=k10_relat_1(X1,k2_relat_1(esk2_0))|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])).
cnf(c_0_29, negated_conjecture, (k10_relat_1(esk2_0,k2_relat_1(esk2_0))=k1_relat_1(esk2_0)), inference(spm,[status(thm)],[c_0_24, c_0_15])).
cnf(c_0_30, negated_conjecture, (k1_funct_1(X1,k1_funct_1(esk2_0,esk1_0))=esk1_0|X1!=k2_funct_1(esk2_0)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_25])).
cnf(c_0_31, negated_conjecture, (v1_funct_1(k2_funct_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_14]), c_0_15])])).
cnf(c_0_32, plain, (k1_funct_1(k5_relat_1(X2,X1),X3)=k1_funct_1(X1,k1_funct_1(X2,X3))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~r2_hidden(X3,k1_relat_1(k5_relat_1(X2,X1)))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, negated_conjecture, (k1_relat_1(k5_relat_1(esk2_0,k2_funct_1(esk2_0)))=k1_relat_1(esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_15]), c_0_29])).
cnf(c_0_34, negated_conjecture, (esk1_0!=k1_funct_1(k2_funct_1(esk2_0),k1_funct_1(esk2_0,esk1_0))|esk1_0!=k1_funct_1(k5_relat_1(esk2_0,k2_funct_1(esk2_0)),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_35, negated_conjecture, (k1_funct_1(k2_funct_1(esk2_0),k1_funct_1(esk2_0,esk1_0))=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_30]), c_0_31]), c_0_22])])).
cnf(c_0_36, negated_conjecture, (k1_funct_1(k5_relat_1(esk2_0,k2_funct_1(esk2_0)),X1)=k1_funct_1(k2_funct_1(esk2_0),k1_funct_1(esk2_0,X1))|~r2_hidden(X1,k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_14]), c_0_31]), c_0_15]), c_0_22])])).
cnf(c_0_37, negated_conjecture, (k1_funct_1(k5_relat_1(esk2_0,k2_funct_1(esk2_0)),esk1_0)!=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_35])])).
cnf(c_0_38, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_20]), c_0_35]), c_0_37]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 39
# Proof object clause steps            : 24
# Proof object formula steps           : 15
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 25
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 26
# Processed clauses                    : 74
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 74
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 148
# ...of the previous two non-trivial   : 148
# Contextual simplify-reflections      : 0
# Paramodulations                      : 143
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 71
#    Positive orientable unit clauses  : 40
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 30
# Current number of unprocessed clauses: 100
# ...number of literals in the above   : 335
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3
# Clause-clause subsumption calls (NU) : 409
# Rec. Clause-clause subsumption calls : 26
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 231
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 103
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6688

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 26068 pages
