# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t41_enumset1, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_xboole_0(k1_tarski(X1),k1_tarski(X2)), file('data/MPTP/proofs/45951618617240b7800e811f05dbb66b.E_input', t41_enumset1)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/45951618617240b7800e811f05dbb66b.E_input', t98_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/45951618617240b7800e811f05dbb66b.E_input', t100_xboole_1)).
fof(t42_enumset1, axiom, ![X1, X2, X3]:k1_enumset1(X1,X2,X3)=k2_xboole_0(k1_tarski(X1),k2_tarski(X2,X3)), file('data/MPTP/proofs/45951618617240b7800e811f05dbb66b.E_input', t42_enumset1)).
fof(t4_xboole_1, axiom, ![X1, X2, X3]:k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/45951618617240b7800e811f05dbb66b.E_input', t4_xboole_1)).
fof(t91_xboole_1, axiom, ![X1, X2, X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3)), file('data/MPTP/proofs/45951618617240b7800e811f05dbb66b.E_input', t91_xboole_1)).
fof(t44_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_tarski(X1),k1_enumset1(X2,X3,X4)), file('data/MPTP/proofs/45951618617240b7800e811f05dbb66b.E_input', t44_enumset1)).
fof(t47_enumset1, axiom, ![X1, X2, X3, X4, X5]:k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k1_tarski(X1),k2_enumset1(X2,X3,X4,X5)), file('data/MPTP/proofs/45951618617240b7800e811f05dbb66b.E_input', t47_enumset1)).
fof(t62_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k1_tarski(X1),k5_enumset1(X2,X3,X4,X5,X6,X7,X8)), file('data/MPTP/proofs/45951618617240b7800e811f05dbb66b.E_input', t62_enumset1)).
fof(t66_enumset1, conjecture, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_enumset1(X6,X7,X8)), file('data/MPTP/proofs/45951618617240b7800e811f05dbb66b.E_input', t66_enumset1)).
fof(l75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8)), file('data/MPTP/proofs/45951618617240b7800e811f05dbb66b.E_input', l75_enumset1)).
fof(c_0_11, plain, ![X45, X46]:k2_tarski(X45,X46)=k2_xboole_0(k1_tarski(X45),k1_tarski(X46)), inference(variable_rename,[status(thm)],[t41_enumset1])).
fof(c_0_12, plain, ![X55, X56]:k2_xboole_0(X55,X56)=k5_xboole_0(X55,k4_xboole_0(X56,X55)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_13, plain, ![X37, X38]:k4_xboole_0(X37,X38)=k5_xboole_0(X37,k3_xboole_0(X37,X38)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_14, plain, ![X42, X43, X44]:k1_enumset1(X42,X43,X44)=k2_xboole_0(k1_tarski(X42),k2_tarski(X43,X44)), inference(variable_rename,[status(thm)],[t42_enumset1])).
cnf(c_0_15, plain, (k2_tarski(X1,X2)=k2_xboole_0(k1_tarski(X1),k1_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_18, plain, ![X17, X18, X19]:k2_xboole_0(k2_xboole_0(X17,X18),X19)=k2_xboole_0(X17,k2_xboole_0(X18,X19)), inference(variable_rename,[status(thm)],[t4_xboole_1])).
cnf(c_0_19, plain, (k1_enumset1(X1,X2,X3)=k2_xboole_0(k1_tarski(X1),k2_tarski(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (k2_tarski(X1,X2)=k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k3_xboole_0(k1_tarski(X2),k1_tarski(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_16]), c_0_17])).
fof(c_0_21, plain, ![X39, X40, X41]:k5_xboole_0(k5_xboole_0(X39,X40),X41)=k5_xboole_0(X39,k5_xboole_0(X40,X41)), inference(variable_rename,[status(thm)],[t91_xboole_1])).
cnf(c_0_22, plain, (k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_23, plain, ![X28, X29, X30, X31]:k2_enumset1(X28,X29,X30,X31)=k2_xboole_0(k1_tarski(X28),k1_enumset1(X29,X30,X31)), inference(variable_rename,[status(thm)],[t44_enumset1])).
cnf(c_0_24, plain, (k1_enumset1(X1,X2,X3)=k5_xboole_0(k1_tarski(X1),k5_xboole_0(k5_xboole_0(k1_tarski(X2),k5_xboole_0(k1_tarski(X3),k3_xboole_0(k1_tarski(X3),k1_tarski(X2)))),k3_xboole_0(k5_xboole_0(k1_tarski(X2),k5_xboole_0(k1_tarski(X3),k3_xboole_0(k1_tarski(X3),k1_tarski(X2)))),k1_tarski(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_16]), c_0_17]), c_0_20]), c_0_20])).
cnf(c_0_25, plain, (k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_26, plain, (k5_xboole_0(k5_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1))),k5_xboole_0(X3,k3_xboole_0(X3,k5_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1))))))=k5_xboole_0(X1,k5_xboole_0(k5_xboole_0(X2,k5_xboole_0(X3,k3_xboole_0(X3,X2))),k3_xboole_0(k5_xboole_0(X2,k5_xboole_0(X3,k3_xboole_0(X3,X2))),X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_16]), c_0_16]), c_0_16]), c_0_16]), c_0_17]), c_0_17]), c_0_17]), c_0_17]), c_0_17])).
fof(c_0_27, plain, ![X32, X33, X34, X35, X36]:k3_enumset1(X32,X33,X34,X35,X36)=k2_xboole_0(k1_tarski(X32),k2_enumset1(X33,X34,X35,X36)), inference(variable_rename,[status(thm)],[t47_enumset1])).
cnf(c_0_28, plain, (k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_tarski(X1),k1_enumset1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k5_xboole_0(k1_tarski(X3),k5_xboole_0(k3_xboole_0(k1_tarski(X3),k1_tarski(X2)),k3_xboole_0(k5_xboole_0(k1_tarski(X2),k5_xboole_0(k1_tarski(X3),k3_xboole_0(k1_tarski(X3),k1_tarski(X2)))),k1_tarski(X1))))))=k1_enumset1(X1,X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25]), c_0_25])).
cnf(c_0_30, plain, (k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(k3_xboole_0(X2,X1),k5_xboole_0(X3,k3_xboole_0(X3,k5_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1))))))))=k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X3,k5_xboole_0(k3_xboole_0(X3,X2),k3_xboole_0(k5_xboole_0(X2,k5_xboole_0(X3,k3_xboole_0(X3,X2))),X1)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_25]), c_0_25]), c_0_25]), c_0_25])).
cnf(c_0_31, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k1_tarski(X1),k2_enumset1(X2,X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_32, plain, (k2_enumset1(X1,X2,X3,X4)=k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_enumset1(X2,X3,X4),k3_xboole_0(k1_enumset1(X2,X3,X4),k1_tarski(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_16]), c_0_17])).
cnf(c_0_33, plain, (k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k5_xboole_0(k3_xboole_0(k1_tarski(X2),k1_tarski(X1)),k5_xboole_0(k1_tarski(X3),k3_xboole_0(k1_tarski(X3),k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k3_xboole_0(k1_tarski(X2),k1_tarski(X1)))))))))=k1_enumset1(X1,X2,X3)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_34, plain, (k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X3,k5_xboole_0(k3_xboole_0(X3,k5_xboole_0(X1,X2)),k5_xboole_0(X4,k3_xboole_0(X4,k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X3,k3_xboole_0(X3,k5_xboole_0(X1,X2)))))))))))=k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X3,k5_xboole_0(X4,k5_xboole_0(k3_xboole_0(X4,X3),k3_xboole_0(k5_xboole_0(X3,k5_xboole_0(X4,k3_xboole_0(X4,X3))),k5_xboole_0(X1,X2)))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_30]), c_0_25]), c_0_25])).
cnf(c_0_35, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k5_xboole_0(k1_tarski(X1),k5_xboole_0(k5_xboole_0(k1_tarski(X2),k5_xboole_0(k1_enumset1(X3,X4,X5),k3_xboole_0(k1_enumset1(X3,X4,X5),k1_tarski(X2)))),k3_xboole_0(k5_xboole_0(k1_tarski(X2),k5_xboole_0(k1_enumset1(X3,X4,X5),k3_xboole_0(k1_enumset1(X3,X4,X5),k1_tarski(X2)))),k1_tarski(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_16]), c_0_17]), c_0_32]), c_0_32])).
cnf(c_0_36, plain, (k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k5_xboole_0(k3_xboole_0(k1_tarski(X2),k1_tarski(X1)),k5_xboole_0(k1_tarski(X3),k5_xboole_0(k3_xboole_0(k1_tarski(X3),k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k3_xboole_0(k1_tarski(X2),k1_tarski(X1))))),X4)))))=k5_xboole_0(k1_enumset1(X1,X2,X3),X4)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_33]), c_0_25]), c_0_25]), c_0_25])).
cnf(c_0_37, plain, (k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X3,k5_xboole_0(X4,k5_xboole_0(k3_xboole_0(X4,k5_xboole_0(X1,k5_xboole_0(X2,X3))),k5_xboole_0(X5,k3_xboole_0(X5,k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X3,k5_xboole_0(X4,k3_xboole_0(X4,k5_xboole_0(X1,k5_xboole_0(X2,X3))))))))))))))=k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X3,k5_xboole_0(X4,k5_xboole_0(X5,k5_xboole_0(k3_xboole_0(X5,X4),k3_xboole_0(k5_xboole_0(X4,k5_xboole_0(X5,k3_xboole_0(X5,X4))),k5_xboole_0(X1,k5_xboole_0(X2,X3)))))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_34]), c_0_25]), c_0_25]), c_0_25]), c_0_25]), c_0_25])).
cnf(c_0_38, plain, (k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(k3_xboole_0(X2,X1),k5_xboole_0(X3,k5_xboole_0(X4,k3_xboole_0(k5_xboole_0(X3,X4),k5_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1)))))))))=k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X3,k5_xboole_0(X4,k5_xboole_0(k3_xboole_0(k5_xboole_0(X3,X4),X2),k3_xboole_0(k5_xboole_0(X2,k5_xboole_0(X3,k5_xboole_0(X4,k3_xboole_0(k5_xboole_0(X3,X4),X2)))),X1))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_25]), c_0_25]), c_0_25])).
cnf(c_0_39, plain, (k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k5_xboole_0(k1_enumset1(X3,X4,X5),k5_xboole_0(k3_xboole_0(k1_enumset1(X3,X4,X5),k1_tarski(X2)),k3_xboole_0(k5_xboole_0(k1_tarski(X2),k5_xboole_0(k1_enumset1(X3,X4,X5),k3_xboole_0(k1_enumset1(X3,X4,X5),k1_tarski(X2)))),k1_tarski(X1))))))=k3_enumset1(X1,X2,X3,X4,X5)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_25]), c_0_25])).
cnf(c_0_40, plain, (k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k5_xboole_0(k3_xboole_0(k1_tarski(X2),k1_tarski(X1)),k5_xboole_0(k1_tarski(X3),k5_xboole_0(X4,k5_xboole_0(k3_xboole_0(X4,k1_tarski(X3)),k3_xboole_0(k5_xboole_0(k1_tarski(X3),k5_xboole_0(X4,k3_xboole_0(X4,k1_tarski(X3)))),k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k3_xboole_0(k1_tarski(X2),k1_tarski(X1)))))))))))=k5_xboole_0(k1_enumset1(X1,X2,X3),k5_xboole_0(X4,k3_xboole_0(X4,k1_enumset1(X1,X2,X3))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_33])).
cnf(c_0_41, plain, (k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(k3_xboole_0(X2,X1),k5_xboole_0(X3,k5_xboole_0(X4,k5_xboole_0(X5,k3_xboole_0(k5_xboole_0(X3,k5_xboole_0(X4,X5)),k5_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1))))))))))=k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X3,k5_xboole_0(X4,k5_xboole_0(X5,k5_xboole_0(k3_xboole_0(k5_xboole_0(X3,k5_xboole_0(X4,X5)),X2),k3_xboole_0(k5_xboole_0(X2,k5_xboole_0(X3,k5_xboole_0(X4,k5_xboole_0(X5,k3_xboole_0(k5_xboole_0(X3,k5_xboole_0(X4,X5)),X2))))),X1)))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_25]), c_0_25]), c_0_25]), c_0_25])).
fof(c_0_42, plain, ![X47, X48, X49, X50, X51, X52, X53, X54]:k6_enumset1(X47,X48,X49,X50,X51,X52,X53,X54)=k2_xboole_0(k1_tarski(X47),k5_enumset1(X48,X49,X50,X51,X52,X53,X54)), inference(variable_rename,[status(thm)],[t62_enumset1])).
fof(c_0_43, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_enumset1(X6,X7,X8))), inference(assume_negation,[status(cth)],[t66_enumset1])).
cnf(c_0_44, plain, (k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k5_xboole_0(k3_xboole_0(k1_tarski(X2),k1_tarski(X1)),k5_xboole_0(k1_enumset1(X3,X4,X5),k3_xboole_0(k1_enumset1(X3,X4,X5),k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k3_xboole_0(k1_tarski(X2),k1_tarski(X1)))))))))=k3_enumset1(X1,X2,X3,X4,X5)), inference(spm,[status(thm)],[c_0_30, c_0_39])).
cnf(c_0_45, plain, (k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k5_xboole_0(k1_tarski(X3),k5_xboole_0(X4,k5_xboole_0(k3_xboole_0(X4,k1_tarski(X3)),k5_xboole_0(k3_xboole_0(k5_xboole_0(k1_tarski(X3),k5_xboole_0(X4,k3_xboole_0(X4,k1_tarski(X3)))),k1_tarski(X2)),k3_xboole_0(k5_xboole_0(k1_tarski(X2),k5_xboole_0(k1_tarski(X3),k5_xboole_0(X4,k5_xboole_0(k3_xboole_0(X4,k1_tarski(X3)),k3_xboole_0(k5_xboole_0(k1_tarski(X3),k5_xboole_0(X4,k3_xboole_0(X4,k1_tarski(X3)))),k1_tarski(X2)))))),k1_tarski(X1))))))))=k5_xboole_0(k1_enumset1(X1,X2,X3),k5_xboole_0(X4,k3_xboole_0(X4,k1_enumset1(X1,X2,X3))))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_46, plain, (k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k5_xboole_0(k1_enumset1(X3,X4,X5),k5_xboole_0(k3_xboole_0(k1_enumset1(X3,X4,X5),k1_tarski(X2)),k5_xboole_0(k3_xboole_0(k5_xboole_0(k1_tarski(X2),k5_xboole_0(k1_enumset1(X3,X4,X5),k3_xboole_0(k1_enumset1(X3,X4,X5),k1_tarski(X2)))),k1_tarski(X1)),X6)))))=k5_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),X6)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_39]), c_0_25]), c_0_25]), c_0_25])).
cnf(c_0_47, plain, (k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k5_xboole_0(k1_tarski(X3),k5_xboole_0(k3_xboole_0(k1_tarski(X3),k1_tarski(X2)),k5_xboole_0(k3_xboole_0(k5_xboole_0(k1_tarski(X2),k5_xboole_0(k1_tarski(X3),k3_xboole_0(k1_tarski(X3),k1_tarski(X2)))),k1_tarski(X1)),X4)))))=k5_xboole_0(k1_enumset1(X1,X2,X3),X4)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_29]), c_0_25]), c_0_25]), c_0_25])).
fof(c_0_48, plain, ![X20, X21, X22, X23, X24, X25, X26, X27]:k6_enumset1(X20,X21,X22,X23,X24,X25,X26,X27)=k2_xboole_0(k2_enumset1(X20,X21,X22,X23),k2_enumset1(X24,X25,X26,X27)), inference(variable_rename,[status(thm)],[l75_enumset1])).
cnf(c_0_49, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k1_tarski(X1),k5_enumset1(X2,X3,X4,X5,X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_50, negated_conjecture, k6_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0)!=k2_xboole_0(k3_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0),k1_enumset1(esk6_0,esk7_0,esk8_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_43])])])).
cnf(c_0_51, plain, (k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k5_xboole_0(k3_xboole_0(k1_tarski(X2),k1_tarski(X1)),k5_xboole_0(k1_enumset1(X3,X4,X5),k5_xboole_0(k3_xboole_0(k1_enumset1(X3,X4,X5),k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k3_xboole_0(k1_tarski(X2),k1_tarski(X1))))),X6)))))=k5_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),X6)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_44]), c_0_25]), c_0_25]), c_0_25])).
cnf(c_0_52, plain, (k5_xboole_0(k1_enumset1(X1,X2,X3),k5_xboole_0(k1_enumset1(X4,X5,X6),k3_xboole_0(k1_enumset1(X4,X5,X6),k1_enumset1(X1,X2,X3))))=k5_xboole_0(k1_tarski(X1),k5_xboole_0(k3_enumset1(X2,X3,X4,X5,X6),k3_xboole_0(k3_enumset1(X2,X3,X4,X5,X6),k1_tarski(X1))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_39]), c_0_46])).
cnf(c_0_53, plain, (k5_xboole_0(X1,k5_xboole_0(k1_tarski(X2),k5_xboole_0(k3_xboole_0(k1_tarski(X2),X1),k5_xboole_0(k1_tarski(X3),k5_xboole_0(k1_tarski(X4),k5_xboole_0(k3_xboole_0(k1_tarski(X4),k1_tarski(X3)),k3_xboole_0(k5_xboole_0(k1_tarski(X3),k5_xboole_0(k1_tarski(X4),k3_xboole_0(k1_tarski(X4),k1_tarski(X3)))),k5_xboole_0(X1,k5_xboole_0(k1_tarski(X2),k3_xboole_0(k1_tarski(X2),X1))))))))))=k5_xboole_0(X1,k5_xboole_0(k1_enumset1(X2,X3,X4),k3_xboole_0(k1_enumset1(X2,X3,X4),X1)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_29]), c_0_47])).
cnf(c_0_54, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_55, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k5_xboole_0(k1_tarski(X1),k5_xboole_0(k5_enumset1(X2,X3,X4,X5,X6,X7,X8),k3_xboole_0(k5_enumset1(X2,X3,X4,X5,X6,X7,X8),k1_tarski(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_16]), c_0_17])).
cnf(c_0_56, negated_conjecture, (k6_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0)!=k2_xboole_0(k3_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0),k1_enumset1(esk6_0,esk7_0,esk8_0))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_57, plain, (k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k5_xboole_0(k3_xboole_0(k1_tarski(X2),k1_tarski(X1)),k5_xboole_0(k1_enumset1(X3,X4,X5),k5_xboole_0(X6,k5_xboole_0(k3_xboole_0(X6,k1_enumset1(X3,X4,X5)),k3_xboole_0(k5_xboole_0(k1_enumset1(X3,X4,X5),k5_xboole_0(X6,k3_xboole_0(X6,k1_enumset1(X3,X4,X5)))),k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k3_xboole_0(k1_tarski(X2),k1_tarski(X1)))))))))))=k5_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k5_xboole_0(X6,k3_xboole_0(X6,k3_enumset1(X1,X2,X3,X4,X5))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_37]), c_0_44])).
cnf(c_0_58, plain, (k5_xboole_0(k1_enumset1(X1,X2,X3),k5_xboole_0(k1_enumset1(X4,X5,X6),k5_xboole_0(k3_xboole_0(k1_enumset1(X4,X5,X6),k1_enumset1(X1,X2,X3)),X7)))=k5_xboole_0(k1_tarski(X1),k5_xboole_0(k3_enumset1(X2,X3,X4,X5,X6),k5_xboole_0(k3_xboole_0(k3_enumset1(X2,X3,X4,X5,X6),k1_tarski(X1)),X7)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_52]), c_0_25]), c_0_25]), c_0_25])).
cnf(c_0_59, plain, (k5_xboole_0(k1_enumset1(X1,X2,X3),k5_xboole_0(k1_tarski(X4),k3_xboole_0(k1_tarski(X4),k1_enumset1(X1,X2,X3))))=k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_enumset1(X2,X3,X4),k3_xboole_0(k1_enumset1(X2,X3,X4),k1_tarski(X1))))), inference(spm,[status(thm)],[c_0_40, c_0_53])).
cnf(c_0_60, plain, (k5_xboole_0(k1_tarski(X1),k5_xboole_0(k5_enumset1(X2,X3,X4,X5,X6,X7,X8),k3_xboole_0(k5_enumset1(X2,X3,X4,X5,X6,X7,X8),k1_tarski(X1))))=k5_xboole_0(k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_enumset1(X2,X3,X4),k3_xboole_0(k1_enumset1(X2,X3,X4),k1_tarski(X1)))),k5_xboole_0(k5_xboole_0(k1_tarski(X5),k5_xboole_0(k1_enumset1(X6,X7,X8),k3_xboole_0(k1_enumset1(X6,X7,X8),k1_tarski(X5)))),k3_xboole_0(k5_xboole_0(k1_tarski(X5),k5_xboole_0(k1_enumset1(X6,X7,X8),k3_xboole_0(k1_enumset1(X6,X7,X8),k1_tarski(X5)))),k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_enumset1(X2,X3,X4),k3_xboole_0(k1_enumset1(X2,X3,X4),k1_tarski(X1)))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_16]), c_0_17]), c_0_32]), c_0_32]), c_0_32]), c_0_32]), c_0_55])).
cnf(c_0_61, negated_conjecture, (k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(k5_enumset1(esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0),k3_xboole_0(k5_enumset1(esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0),k1_tarski(esk1_0))))!=k5_xboole_0(k3_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0),k5_xboole_0(k1_enumset1(esk6_0,esk7_0,esk8_0),k3_xboole_0(k1_enumset1(esk6_0,esk7_0,esk8_0),k3_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_16]), c_0_17]), c_0_55])).
cnf(c_0_62, plain, (k5_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k5_xboole_0(k1_enumset1(X6,X7,X8),k3_xboole_0(k1_enumset1(X6,X7,X8),k3_enumset1(X1,X2,X3,X4,X5))))=k5_xboole_0(k1_enumset1(X1,X2,X3),k5_xboole_0(k3_enumset1(X4,X5,X6,X7,X8),k3_xboole_0(k3_enumset1(X4,X5,X6,X7,X8),k1_enumset1(X1,X2,X3))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_52]), c_0_58]), c_0_40])).
cnf(c_0_63, plain, (k5_xboole_0(k1_enumset1(X1,X2,X3),k5_xboole_0(k1_tarski(X4),k5_xboole_0(k3_xboole_0(k1_tarski(X4),k1_enumset1(X1,X2,X3)),X5)))=k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_enumset1(X2,X3,X4),k5_xboole_0(k3_xboole_0(k1_enumset1(X2,X3,X4),k1_tarski(X1)),X5)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_59]), c_0_25]), c_0_25]), c_0_25])).
cnf(c_0_64, plain, (k5_xboole_0(X1,k5_xboole_0(k1_tarski(X2),k5_xboole_0(k3_xboole_0(k1_tarski(X2),X1),k5_xboole_0(k1_tarski(X3),k5_xboole_0(k1_enumset1(X4,X5,X6),k5_xboole_0(k3_xboole_0(k1_enumset1(X4,X5,X6),k1_tarski(X3)),k3_xboole_0(k5_xboole_0(k1_tarski(X3),k5_xboole_0(k1_enumset1(X4,X5,X6),k3_xboole_0(k1_enumset1(X4,X5,X6),k1_tarski(X3)))),k5_xboole_0(X1,k5_xboole_0(k1_tarski(X2),k3_xboole_0(k1_tarski(X2),X1))))))))))=k5_xboole_0(X1,k5_xboole_0(k3_enumset1(X2,X3,X4,X5,X6),k3_xboole_0(k3_enumset1(X2,X3,X4,X5,X6),X1)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_39]), c_0_46])).
cnf(c_0_65, plain, (k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_enumset1(X2,X3,X4),k5_xboole_0(k3_xboole_0(k1_enumset1(X2,X3,X4),k1_tarski(X1)),k5_xboole_0(k1_tarski(X5),k5_xboole_0(k1_enumset1(X6,X7,X8),k5_xboole_0(k3_xboole_0(k1_enumset1(X6,X7,X8),k1_tarski(X5)),k3_xboole_0(k5_xboole_0(k1_tarski(X5),k5_xboole_0(k1_enumset1(X6,X7,X8),k3_xboole_0(k1_enumset1(X6,X7,X8),k1_tarski(X5)))),k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_enumset1(X2,X3,X4),k3_xboole_0(k1_enumset1(X2,X3,X4),k1_tarski(X1)))))))))))=k5_xboole_0(k1_tarski(X1),k5_xboole_0(k5_enumset1(X2,X3,X4,X5,X6,X7,X8),k3_xboole_0(k5_enumset1(X2,X3,X4,X5,X6,X7,X8),k1_tarski(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_25]), c_0_25]), c_0_25]), c_0_25])).
cnf(c_0_66, negated_conjecture, (k5_xboole_0(k1_enumset1(esk1_0,esk2_0,esk3_0),k5_xboole_0(k3_enumset1(esk4_0,esk5_0,esk6_0,esk7_0,esk8_0),k3_xboole_0(k3_enumset1(esk4_0,esk5_0,esk6_0,esk7_0,esk8_0),k1_enumset1(esk1_0,esk2_0,esk3_0))))!=k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(k5_enumset1(esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0),k3_xboole_0(k5_enumset1(esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0),k1_tarski(esk1_0))))), inference(rw,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_67, plain, (k5_xboole_0(k1_enumset1(X1,X2,X3),k5_xboole_0(k3_enumset1(X4,X5,X6,X7,X8),k3_xboole_0(k3_enumset1(X4,X5,X6,X7,X8),k1_enumset1(X1,X2,X3))))=k5_xboole_0(k1_tarski(X1),k5_xboole_0(k5_enumset1(X2,X3,X4,X5,X6,X7,X8),k3_xboole_0(k5_enumset1(X2,X3,X4,X5,X6,X7,X8),k1_tarski(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_59]), c_0_65])).
cnf(c_0_68, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_67])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 69
# Proof object clause steps            : 46
# Proof object formula steps           : 23
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 87
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 11
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 6
# Processed clauses                    : 153
# ...of these trivial                  : 11
# ...subsumed                          : 82
# ...remaining for further processing  : 60
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 4157
# ...of the previous two non-trivial   : 4040
# Contextual simplify-reflections      : 0
# Paramodulations                      : 4157
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 51
#    Positive orientable unit clauses  : 33
#    Positive unorientable unit clauses: 18
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 3898
# ...number of literals in the above   : 3898
# Current number of archived formulas  : 0
# Current number of archived clauses   : 14
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2119
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 440920

# -------------------------------------------------
# User time                : 0.137 s
# System time              : 0.003 s
# Total time               : 0.140 s
# Maximum resident set size: 28600 pages
