# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t85_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_xboole_0(X1,k4_xboole_0(X3,X2))), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', t85_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', t100_xboole_1)).
fof(t31_subset_1, conjecture, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(r1_tarski(X2,X3)<=>r1_tarski(k3_subset_1(X1,X3),k3_subset_1(X1,X2))))), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', t31_subset_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', t36_xboole_1)).
fof(t44_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k4_xboole_0(X1,X2),X3)=>r1_tarski(X1,k2_xboole_0(X2,X3))), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', t44_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', t28_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', d10_xboole_0)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', d1_zfmisc_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', d2_subset_1)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', fc1_subset_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', commutativity_k3_xboole_0)).
fof(t86_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_xboole_0(X1,X3))=>r1_tarski(X1,k4_xboole_0(X2,X3))), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', t86_xboole_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', t7_xboole_1)).
fof(t73_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,k2_xboole_0(X2,X3))&r1_xboole_0(X1,X3))=>r1_tarski(X1,X2)), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', t73_xboole_1)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', t1_xboole_1)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', symmetry_r1_xboole_0)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', t17_xboole_1)).
fof(t79_xboole_1, axiom, ![X1, X2]:r1_xboole_0(k4_xboole_0(X1,X2),X2), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', t79_xboole_1)).
fof(t11_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k2_xboole_0(X1,X2),X3)=>r1_tarski(X1,X3)), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', t11_xboole_1)).
fof(l22_zfmisc_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>k2_xboole_0(k1_tarski(X1),X2)=X2), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', l22_zfmisc_1)).
fof(l1_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', l1_zfmisc_1)).
fof(d5_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', d5_subset_1)).
fof(t79_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>r1_tarski(k1_zfmisc_1(X1),k1_zfmisc_1(X2))), file('data/MPTP/proofs/46087ec39a254e85a08f13ee765e037b.E_input', t79_zfmisc_1)).
fof(c_0_23, plain, ![X51, X52, X53]:(~r1_tarski(X51,X52)|r1_xboole_0(X51,k4_xboole_0(X53,X52))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t85_xboole_1])])).
fof(c_0_24, plain, ![X49, X50]:k4_xboole_0(X49,X50)=k5_xboole_0(X49,k3_xboole_0(X49,X50)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_25, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(r1_tarski(X2,X3)<=>r1_tarski(k3_subset_1(X1,X3),k3_subset_1(X1,X2)))))), inference(assume_negation,[status(cth)],[t31_subset_1])).
fof(c_0_26, plain, ![X47, X48]:r1_tarski(k4_xboole_0(X47,X48),X47), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_27, plain, ![X56, X57, X58]:(~r1_tarski(k4_xboole_0(X56,X57),X58)|r1_tarski(X56,k2_xboole_0(X57,X58))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t44_xboole_1])])).
cnf(c_0_28, plain, (r1_xboole_0(X1,k4_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_30, plain, ![X27, X28]:(~r1_tarski(X27,X28)|k3_xboole_0(X27,X28)=X27), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_31, plain, ![X59, X60]:(((r1_tarski(X59,X60)|X59!=X60)&(r1_tarski(X60,X59)|X59!=X60))&(~r1_tarski(X59,X60)|~r1_tarski(X60,X59)|X59=X60)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_32, plain, ![X20, X21, X22, X23, X24, X25]:(((~r2_hidden(X22,X21)|r1_tarski(X22,X20)|X21!=k1_zfmisc_1(X20))&(~r1_tarski(X23,X20)|r2_hidden(X23,X21)|X21!=k1_zfmisc_1(X20)))&((~r2_hidden(esk4_2(X24,X25),X25)|~r1_tarski(esk4_2(X24,X25),X24)|X25=k1_zfmisc_1(X24))&(r2_hidden(esk4_2(X24,X25),X25)|r1_tarski(esk4_2(X24,X25),X24)|X25=k1_zfmisc_1(X24)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
fof(c_0_33, plain, ![X45, X46]:(((~m1_subset_1(X46,X45)|r2_hidden(X46,X45)|v1_xboole_0(X45))&(~r2_hidden(X46,X45)|m1_subset_1(X46,X45)|v1_xboole_0(X45)))&((~m1_subset_1(X46,X45)|v1_xboole_0(X46)|~v1_xboole_0(X45))&(~v1_xboole_0(X46)|m1_subset_1(X46,X45)|~v1_xboole_0(X45)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_34, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))&((~r1_tarski(esk2_0,esk3_0)|~r1_tarski(k3_subset_1(esk1_0,esk3_0),k3_subset_1(esk1_0,esk2_0)))&(r1_tarski(esk2_0,esk3_0)|r1_tarski(k3_subset_1(esk1_0,esk3_0),k3_subset_1(esk1_0,esk2_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])).
fof(c_0_35, plain, ![X10]:~v1_xboole_0(k1_zfmisc_1(X10)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
cnf(c_0_36, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_37, plain, ![X14, X15]:k3_xboole_0(X14,X15)=k3_xboole_0(X15,X14), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_38, plain, ![X38, X39, X40]:(~r1_tarski(X38,X39)|~r1_xboole_0(X38,X40)|r1_tarski(X38,k4_xboole_0(X39,X40))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t86_xboole_1])])).
cnf(c_0_39, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(k4_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_40, plain, (r1_xboole_0(X1,k5_xboole_0(X3,k3_xboole_0(X3,X2)))|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_41, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_42, plain, ![X54, X55]:r1_tarski(X54,k2_xboole_0(X54,X55)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_43, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_44, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_45, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_47, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_48, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_49, plain, (r1_tarski(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X1)), inference(rw,[status(thm)],[c_0_36, c_0_29])).
cnf(c_0_50, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_51, plain, (r1_tarski(X1,k4_xboole_0(X2,X3))|~r1_tarski(X1,X2)|~r1_xboole_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_52, plain, ![X7, X8, X9]:(~r1_tarski(X7,k2_xboole_0(X8,X9))|~r1_xboole_0(X7,X9)|r1_tarski(X7,X8)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t73_xboole_1])])).
cnf(c_0_53, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X3)), inference(rw,[status(thm)],[c_0_39, c_0_29])).
cnf(c_0_54, plain, (r1_xboole_0(X1,k5_xboole_0(X2,X2))|~r1_tarski(X1,X3)|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_55, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_56, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_43])).
fof(c_0_57, plain, ![X11, X12, X13]:(~r1_tarski(X11,X12)|~r1_tarski(X12,X13)|r1_tarski(X11,X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
cnf(c_0_58, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_59, negated_conjecture, (r2_hidden(esk3_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])).
fof(c_0_60, plain, ![X41, X42]:(~r1_xboole_0(X41,X42)|r1_xboole_0(X42,X41)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
cnf(c_0_61, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=X1|~r1_tarski(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_62, plain, (k3_xboole_0(X1,X2)=X2|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_41, c_0_50])).
cnf(c_0_63, plain, (r1_tarski(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|~r1_tarski(X1,X2)|~r1_xboole_0(X1,X3)), inference(rw,[status(thm)],[c_0_51, c_0_29])).
cnf(c_0_64, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_xboole_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_65, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_53, c_0_49])).
cnf(c_0_66, plain, (r1_xboole_0(X1,k5_xboole_0(X2,X2))|~r1_tarski(X2,k2_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_67, plain, (r1_tarski(X1,k2_xboole_0(X2,k5_xboole_0(X1,k3_xboole_0(X1,X2))))), inference(spm,[status(thm)],[c_0_53, c_0_56])).
cnf(c_0_68, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_69, negated_conjecture, (r1_tarski(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_70, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_71, plain, (k5_xboole_0(X1,X2)=X1|~r1_tarski(X1,k5_xboole_0(X1,X2))|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_72, plain, (r1_tarski(X1,k5_xboole_0(X2,X3))|~r1_xboole_0(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_63, c_0_62])).
cnf(c_0_73, plain, (r1_tarski(X1,X2)|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_74, plain, (r1_xboole_0(X1,k5_xboole_0(X2,X2))), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_75, negated_conjecture, (r1_tarski(X1,esk1_0)|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
fof(c_0_76, plain, ![X31, X32]:r1_tarski(k3_xboole_0(X31,X32),X31), inference(variable_rename,[status(thm)],[t17_xboole_1])).
cnf(c_0_77, plain, (r1_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X3)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_70, c_0_40])).
cnf(c_0_78, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=X1|~r1_xboole_0(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_63]), c_0_56])])).
cnf(c_0_79, plain, (k5_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)|~r1_tarski(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_56])])).
cnf(c_0_80, plain, (r1_tarski(k5_xboole_0(X1,X1),X2)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_81, negated_conjecture, (r1_tarski(k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,X1)),esk1_0)), inference(spm,[status(thm)],[c_0_75, c_0_49])).
cnf(c_0_82, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
fof(c_0_83, plain, ![X29, X30]:r1_xboole_0(k4_xboole_0(X29,X30),X30), inference(variable_rename,[status(thm)],[t79_xboole_1])).
cnf(c_0_84, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X1,X3)|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_85, plain, (r1_xboole_0(X1,k5_xboole_0(X2,X3))|~r1_tarski(X1,X3)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_40, c_0_62])).
cnf(c_0_86, plain, (k5_xboole_0(X1,k5_xboole_0(X2,X2))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_74]), c_0_80])])).
cnf(c_0_87, plain, (r1_tarski(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X3)|~r1_tarski(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_73, c_0_40])).
cnf(c_0_88, negated_conjecture, (r1_tarski(k5_xboole_0(esk3_0,X1),esk1_0)|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_81, c_0_62])).
cnf(c_0_89, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_82, c_0_50])).
cnf(c_0_90, plain, (r1_xboole_0(k4_xboole_0(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_91, negated_conjecture, (r1_xboole_0(X1,esk3_0)|~r1_xboole_0(X1,esk1_0)), inference(spm,[status(thm)],[c_0_84, c_0_69])).
cnf(c_0_92, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k5_xboole_0(X3,X3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_80])])).
cnf(c_0_93, negated_conjecture, (r1_tarski(k5_xboole_0(esk3_0,k3_xboole_0(esk1_0,esk3_0)),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_50]), c_0_50]), c_0_89])])).
cnf(c_0_94, plain, (r1_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X2)), inference(rw,[status(thm)],[c_0_90, c_0_29])).
cnf(c_0_95, plain, (r1_tarski(X1,X2)|~r1_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_64, c_0_67])).
cnf(c_0_96, negated_conjecture, (r1_xboole_0(esk3_0,X1)|~r1_xboole_0(X1,esk1_0)), inference(spm,[status(thm)],[c_0_70, c_0_91])).
cnf(c_0_97, negated_conjecture, (r1_xboole_0(k5_xboole_0(esk3_0,k3_xboole_0(esk1_0,esk3_0)),X1)), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_98, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_99, plain, (r1_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1)))), inference(spm,[status(thm)],[c_0_70, c_0_94])).
cnf(c_0_100, plain, (r1_tarski(X1,X2)|~r1_xboole_0(X1,k5_xboole_0(X1,X2))|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_95, c_0_62])).
cnf(c_0_101, negated_conjecture, (r1_xboole_0(esk3_0,k5_xboole_0(esk3_0,k3_xboole_0(esk1_0,esk3_0)))), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_102, negated_conjecture, (r2_hidden(esk2_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_98]), c_0_47])).
cnf(c_0_103, plain, (X1=k5_xboole_0(X2,X2)|~r1_tarski(X1,k5_xboole_0(X2,X2))), inference(spm,[status(thm)],[c_0_48, c_0_80])).
cnf(c_0_104, plain, (r1_xboole_0(X1,k5_xboole_0(X2,X1))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_99, c_0_62])).
cnf(c_0_105, negated_conjecture, (r1_tarski(esk3_0,k3_xboole_0(esk1_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_89])])).
cnf(c_0_106, negated_conjecture, (r1_tarski(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_58, c_0_102])).
cnf(c_0_107, plain, (X1=k5_xboole_0(X2,X2)|~r1_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_72]), c_0_56])])).
cnf(c_0_108, plain, (r1_xboole_0(k3_xboole_0(X1,X2),X1)|~r1_xboole_0(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_78]), c_0_82])])).
cnf(c_0_109, negated_conjecture, (k3_xboole_0(esk1_0,esk3_0)=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_105]), c_0_89])])).
cnf(c_0_110, negated_conjecture, (r1_tarski(X1,esk1_0)|~r1_tarski(X1,esk2_0)), inference(spm,[status(thm)],[c_0_68, c_0_106])).
cnf(c_0_111, plain, (k3_xboole_0(X1,X2)=k5_xboole_0(X1,X1)|~r1_xboole_0(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_82])])).
cnf(c_0_112, negated_conjecture, (r1_xboole_0(esk3_0,k5_xboole_0(esk1_0,esk3_0))), inference(spm,[status(thm)],[c_0_99, c_0_109])).
cnf(c_0_113, negated_conjecture, (r1_tarski(k5_xboole_0(esk2_0,k3_xboole_0(esk2_0,X1)),esk1_0)), inference(spm,[status(thm)],[c_0_110, c_0_49])).
fof(c_0_114, plain, ![X35, X36, X37]:(~r1_tarski(k2_xboole_0(X35,X36),X37)|r1_tarski(X35,X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_xboole_1])])).
fof(c_0_115, plain, ![X16, X17]:(~r2_hidden(X16,X17)|k2_xboole_0(k1_tarski(X16),X17)=X17), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l22_zfmisc_1])])).
cnf(c_0_116, negated_conjecture, (k3_xboole_0(esk3_0,k5_xboole_0(esk1_0,esk3_0))=k5_xboole_0(esk3_0,esk3_0)), inference(spm,[status(thm)],[c_0_111, c_0_112])).
cnf(c_0_117, negated_conjecture, (r1_tarski(k5_xboole_0(esk2_0,X1),esk1_0)|~r1_tarski(X1,esk2_0)), inference(spm,[status(thm)],[c_0_113, c_0_62])).
cnf(c_0_118, plain, (r1_tarski(X1,X3)|~r1_tarski(k2_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_114])).
cnf(c_0_119, plain, (k2_xboole_0(k1_tarski(X1),X2)=X2|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_115])).
cnf(c_0_120, negated_conjecture, (r1_xboole_0(X1,esk3_0)|~r1_tarski(X1,k5_xboole_0(esk1_0,esk3_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_116]), c_0_86])).
cnf(c_0_121, negated_conjecture, (r1_tarski(k5_xboole_0(esk2_0,k3_xboole_0(esk1_0,esk2_0)),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_117]), c_0_50]), c_0_50]), c_0_89])])).
fof(c_0_122, plain, ![X43, X44]:((~r1_tarski(k1_tarski(X43),X44)|r2_hidden(X43,X44))&(~r2_hidden(X43,X44)|r1_tarski(k1_tarski(X43),X44))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])).
cnf(c_0_123, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_118, c_0_119])).
fof(c_0_124, plain, ![X18, X19]:(~m1_subset_1(X19,k1_zfmisc_1(X18))|k3_subset_1(X18,X19)=k4_xboole_0(X18,X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
cnf(c_0_125, negated_conjecture, (r1_xboole_0(k5_xboole_0(esk2_0,k3_xboole_0(esk1_0,esk2_0)),esk3_0)), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_126, plain, (r2_hidden(X1,X2)|~r1_tarski(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_127, negated_conjecture, (r1_tarski(k1_tarski(esk2_0),X1)|~r1_tarski(k1_zfmisc_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_123, c_0_102])).
cnf(c_0_128, plain, (k3_subset_1(X2,X1)=k4_xboole_0(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_129, negated_conjecture, (k5_xboole_0(esk2_0,k3_xboole_0(esk1_0,esk2_0))=k5_xboole_0(esk3_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_125]), c_0_121])])).
cnf(c_0_130, negated_conjecture, (r2_hidden(esk2_0,X1)|~r1_tarski(k1_zfmisc_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_126, c_0_127])).
fof(c_0_131, plain, ![X33, X34]:(~r1_tarski(X33,X34)|r1_tarski(k1_zfmisc_1(X33),k1_zfmisc_1(X34))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t79_zfmisc_1])])).
cnf(c_0_132, plain, (k5_xboole_0(X2,k3_xboole_0(X2,X1))=k3_subset_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_128, c_0_29])).
cnf(c_0_133, negated_conjecture, (r1_tarski(esk2_0,k3_xboole_0(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_129]), c_0_74]), c_0_89])])).
cnf(c_0_134, negated_conjecture, (r1_tarski(esk2_0,X1)|~r1_tarski(k1_zfmisc_1(esk1_0),k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_58, c_0_130])).
cnf(c_0_135, plain, (r1_tarski(k1_zfmisc_1(X1),k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_131])).
cnf(c_0_136, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(k3_subset_1(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_53, c_0_132])).
cnf(c_0_137, negated_conjecture, (r1_tarski(esk2_0,esk3_0)|r1_tarski(k3_subset_1(esk1_0,esk3_0),k3_subset_1(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_138, negated_conjecture, (k3_xboole_0(esk1_0,esk2_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_133]), c_0_89])])).
cnf(c_0_139, negated_conjecture, (~r1_tarski(esk2_0,esk3_0)|~r1_tarski(k3_subset_1(esk1_0,esk3_0),k3_subset_1(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_140, plain, (k3_subset_1(X1,X2)=k5_xboole_0(X1,X2)|~r1_tarski(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_132, c_0_62])).
cnf(c_0_141, negated_conjecture, (r1_tarski(esk2_0,X1)|~r1_tarski(esk1_0,X1)), inference(spm,[status(thm)],[c_0_134, c_0_135])).
cnf(c_0_142, negated_conjecture, (r1_tarski(esk1_0,k2_xboole_0(esk3_0,k3_subset_1(esk1_0,esk2_0)))|r1_tarski(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_137]), c_0_46])])).
cnf(c_0_143, negated_conjecture, (k3_subset_1(esk1_0,esk2_0)=k5_xboole_0(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_138]), c_0_98])])).
cnf(c_0_144, negated_conjecture, (~r1_tarski(k5_xboole_0(esk1_0,esk3_0),k3_subset_1(esk1_0,esk2_0))|~r1_tarski(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140]), c_0_69]), c_0_46])])).
cnf(c_0_145, negated_conjecture, (r1_tarski(esk2_0,X1)|~r1_xboole_0(esk2_0,X2)|~r1_tarski(esk1_0,k2_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_64, c_0_141])).
cnf(c_0_146, negated_conjecture, (r1_tarski(esk1_0,k2_xboole_0(esk3_0,k5_xboole_0(esk1_0,esk2_0)))|r1_tarski(esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_142, c_0_143])).
cnf(c_0_147, negated_conjecture, (r1_xboole_0(esk2_0,k5_xboole_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_99, c_0_138])).
cnf(c_0_148, negated_conjecture, (~r1_tarski(k5_xboole_0(esk1_0,esk3_0),k5_xboole_0(esk1_0,esk2_0))|~r1_tarski(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_140]), c_0_106]), c_0_98])])).
cnf(c_0_149, negated_conjecture, (r1_tarski(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146]), c_0_147])])).
cnf(c_0_150, negated_conjecture, (~r1_tarski(k5_xboole_0(esk1_0,esk3_0),k5_xboole_0(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_148, c_0_149])])).
cnf(c_0_151, negated_conjecture, (r1_tarski(k5_xboole_0(esk1_0,esk3_0),esk1_0)), inference(spm,[status(thm)],[c_0_49, c_0_109])).
cnf(c_0_152, negated_conjecture, (r1_xboole_0(X1,esk2_0)|~r1_xboole_0(X1,esk3_0)), inference(spm,[status(thm)],[c_0_84, c_0_149])).
cnf(c_0_153, negated_conjecture, (r1_xboole_0(k5_xboole_0(esk1_0,esk3_0),esk3_0)), inference(spm,[status(thm)],[c_0_94, c_0_109])).
cnf(c_0_154, negated_conjecture, (~r1_xboole_0(k5_xboole_0(esk1_0,esk3_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_72]), c_0_151]), c_0_106])])).
cnf(c_0_155, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_153]), c_0_154]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 156
# Proof object clause steps            : 109
# Proof object formula steps           : 47
# Proof object conjectures             : 50
# Proof object clause conjectures      : 47
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 23
# Proof object generating inferences   : 72
# Proof object simplifying inferences  : 63
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 23
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 34
# Processed clauses                    : 8779
# ...of these trivial                  : 79
# ...subsumed                          : 6817
# ...remaining for further processing  : 1882
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 176
# Backward-rewritten                   : 98
# Generated clauses                    : 58938
# ...of the previous two non-trivial   : 54563
# Contextual simplify-reflections      : 75
# Paramodulations                      : 58932
# Factorizations                       : 2
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1571
#    Positive orientable unit clauses  : 122
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 1441
# Current number of unprocessed clauses: 45349
# ...number of literals in the above   : 190256
# Current number of archived formulas  : 0
# Current number of archived clauses   : 308
# Clause-clause subsumption calls (NU) : 1071436
# Rec. Clause-clause subsumption calls : 703982
# Non-unit clause-clause subsumptions  : 5952
# Unit Clause-clause subsumption calls : 6004
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 276
# BW rewrite match successes           : 30
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 894658

# -------------------------------------------------
# User time                : 0.735 s
# System time              : 0.026 s
# Total time               : 0.762 s
# Maximum resident set size: 37496 pages
