# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.005 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/467a5a3cc9cf4cb8a1c09bce06541856.E_input', d3_tarski)).
fof(d15_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>(v3_relat_1(X1)<=>r1_tarski(k2_relat_1(X1),k1_tarski(k1_xboole_0)))), file('data/MPTP/proofs/467a5a3cc9cf4cb8a1c09bce06541856.E_input', d15_relat_1)).
fof(t184_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>(v3_relat_1(X1)<=>![X2]:(r2_hidden(X2,k2_relat_1(X1))=>X2=k1_xboole_0))), file('data/MPTP/proofs/467a5a3cc9cf4cb8a1c09bce06541856.E_input', t184_relat_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/467a5a3cc9cf4cb8a1c09bce06541856.E_input', d1_tarski)).
fof(c_0_4, plain, ![X4, X5, X6, X7, X8]:((~r1_tarski(X4,X5)|(~r2_hidden(X6,X4)|r2_hidden(X6,X5)))&((r2_hidden(esk1_2(X7,X8),X7)|r1_tarski(X7,X8))&(~r2_hidden(esk1_2(X7,X8),X8)|r1_tarski(X7,X8)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_5, plain, ![X10]:((~v3_relat_1(X10)|r1_tarski(k2_relat_1(X10),k1_tarski(k1_xboole_0))|~v1_relat_1(X10))&(~r1_tarski(k2_relat_1(X10),k1_tarski(k1_xboole_0))|v3_relat_1(X10)|~v1_relat_1(X10))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d15_relat_1])])])).
fof(c_0_6, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>(v3_relat_1(X1)<=>![X2]:(r2_hidden(X2,k2_relat_1(X1))=>X2=k1_xboole_0)))), inference(assume_negation,[status(cth)],[t184_relat_1])).
fof(c_0_7, plain, ![X14, X15, X16, X17, X18, X19]:(((~r2_hidden(X16,X15)|X16=X14|X15!=k1_tarski(X14))&(X17!=X14|r2_hidden(X17,X15)|X15!=k1_tarski(X14)))&((~r2_hidden(esk4_2(X18,X19),X19)|esk4_2(X18,X19)!=X18|X19=k1_tarski(X18))&(r2_hidden(esk4_2(X18,X19),X19)|esk4_2(X18,X19)=X18|X19=k1_tarski(X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_8, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9, plain, (r1_tarski(k2_relat_1(X1),k1_tarski(k1_xboole_0))|~v3_relat_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_10, negated_conjecture, ![X13]:(v1_relat_1(esk2_0)&(((r2_hidden(esk3_0,k2_relat_1(esk2_0))|~v3_relat_1(esk2_0))&(esk3_0!=k1_xboole_0|~v3_relat_1(esk2_0)))&(v3_relat_1(esk2_0)|(~r2_hidden(X13,k2_relat_1(esk2_0))|X13=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])).
cnf(c_0_11, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (r2_hidden(X1,k1_tarski(k1_xboole_0))|~v3_relat_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k2_relat_1(X2))), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_13, negated_conjecture, (r2_hidden(esk3_0,k2_relat_1(esk2_0))|~v3_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, plain, (X1=X2|~r2_hidden(X1,k1_tarski(X2))), inference(er,[status(thm)],[c_0_11])).
cnf(c_0_16, negated_conjecture, (r2_hidden(esk3_0,k1_tarski(k1_xboole_0))|~v3_relat_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14])])).
cnf(c_0_17, negated_conjecture, (esk3_0!=k1_xboole_0|~v3_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, negated_conjecture, (v3_relat_1(esk2_0)|X1=k1_xboole_0|~r2_hidden(X1,k2_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19, plain, (r2_hidden(esk1_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_20, negated_conjecture, (~v3_relat_1(esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])).
cnf(c_0_21, plain, (r1_tarski(X1,X2)|~r2_hidden(esk1_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_22, negated_conjecture, (esk1_2(k2_relat_1(esk2_0),X1)=k1_xboole_0|r1_tarski(k2_relat_1(esk2_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])).
cnf(c_0_23, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_24, plain, (v3_relat_1(X1)|~r1_tarski(k2_relat_1(X1),k1_tarski(k1_xboole_0))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_25, negated_conjecture, (r1_tarski(k2_relat_1(esk2_0),X1)|~r2_hidden(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_26, plain, (r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_23])])).
cnf(c_0_27, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_14]), c_0_26])]), c_0_20]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 28
# Proof object clause steps            : 19
# Proof object formula steps           : 9
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 13
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 13
# Processed clauses                    : 40
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 38
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 31
# ...of the previous two non-trivial   : 22
# Contextual simplify-reflections      : 1
# Paramodulations                      : 24
# Factorizations                       : 2
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 23
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 19
# Current number of unprocessed clauses: 7
# ...number of literals in the above   : 23
# Current number of archived formulas  : 0
# Current number of archived clauses   : 13
# Clause-clause subsumption calls (NU) : 18
# Rec. Clause-clause subsumption calls : 18
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 6
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1233

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.000 s
# Total time               : 0.006 s
# Maximum resident set size: 26292 pages
