# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t48_orders_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>~(r2_hidden(X2,k1_orders_2(X1,k6_domain_1(u1_struct_0(X1),X2)))))), file('data/MPTP/proofs/46ceff596c6c46beb6dd3f58afb33a41.E_input', t48_orders_2)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/46ceff596c6c46beb6dd3f58afb33a41.E_input', t24_orders_2)).
fof(t38_orders_2, axiom, ![X1]:((v3_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(~((?[X4]:((v6_orders_2(X4,X1)&m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1))))&(r2_hidden(X2,X4)&r2_hidden(X3,X4)))&(~(r1_orders_2(X1,X2,X3))&~(r1_orders_2(X1,X3,X2)))))&~(((r1_orders_2(X1,X2,X3)|r1_orders_2(X1,X3,X2))&![X4]:((v6_orders_2(X4,X1)&m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1))))=>~((r2_hidden(X2,X4)&r2_hidden(X3,X4)))))))))), file('data/MPTP/proofs/46ceff596c6c46beb6dd3f58afb33a41.E_input', t38_orders_2)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/46ceff596c6c46beb6dd3f58afb33a41.E_input', redefinition_k6_domain_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/46ceff596c6c46beb6dd3f58afb33a41.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/46ceff596c6c46beb6dd3f58afb33a41.E_input', t70_enumset1)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/46ceff596c6c46beb6dd3f58afb33a41.E_input', t5_subset)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/46ceff596c6c46beb6dd3f58afb33a41.E_input', d2_tarski)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/46ceff596c6c46beb6dd3f58afb33a41.E_input', t2_subset)).
fof(t47_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>~((r2_hidden(X2,X3)&r2_hidden(X2,k1_orders_2(X1,X3))))))), file('data/MPTP/proofs/46ceff596c6c46beb6dd3f58afb33a41.E_input', t47_orders_2)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/46ceff596c6c46beb6dd3f58afb33a41.E_input', t4_subset)).
fof(t63_subset_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))), file('data/MPTP/proofs/46ceff596c6c46beb6dd3f58afb33a41.E_input', t63_subset_1)).
fof(c_0_12, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>~(r2_hidden(X2,k1_orders_2(X1,k6_domain_1(u1_struct_0(X1),X2))))))), inference(assume_negation,[status(cth)],[t48_orders_2])).
fof(c_0_13, plain, ![X7, X8]:(v2_struct_0(X7)|(~v3_orders_2(X7)|~l1_orders_2(X7))|(~m1_subset_1(X8,u1_struct_0(X7))|r1_orders_2(X7,X8,X8))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_14, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&r2_hidden(esk2_0,k1_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])).
fof(c_0_15, plain, ![X9, X10, X11, X12]:((~v6_orders_2(X12,X9)|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X9)))|(~r2_hidden(X10,X12)|~r2_hidden(X11,X12))|(r1_orders_2(X9,X10,X11)|r1_orders_2(X9,X11,X10))|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(~v3_orders_2(X9)|~l1_orders_2(X9)))&((((v6_orders_2(esk3_3(X9,X10,X11),X9)|~r1_orders_2(X9,X10,X11)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(~v3_orders_2(X9)|~l1_orders_2(X9)))&(m1_subset_1(esk3_3(X9,X10,X11),k1_zfmisc_1(u1_struct_0(X9)))|~r1_orders_2(X9,X10,X11)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(~v3_orders_2(X9)|~l1_orders_2(X9))))&((r2_hidden(X10,esk3_3(X9,X10,X11))|~r1_orders_2(X9,X10,X11)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(~v3_orders_2(X9)|~l1_orders_2(X9)))&(r2_hidden(X11,esk3_3(X9,X10,X11))|~r1_orders_2(X9,X10,X11)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(~v3_orders_2(X9)|~l1_orders_2(X9)))))&(((v6_orders_2(esk3_3(X9,X10,X11),X9)|~r1_orders_2(X9,X11,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(~v3_orders_2(X9)|~l1_orders_2(X9)))&(m1_subset_1(esk3_3(X9,X10,X11),k1_zfmisc_1(u1_struct_0(X9)))|~r1_orders_2(X9,X11,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(~v3_orders_2(X9)|~l1_orders_2(X9))))&((r2_hidden(X10,esk3_3(X9,X10,X11))|~r1_orders_2(X9,X11,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(~v3_orders_2(X9)|~l1_orders_2(X9)))&(r2_hidden(X11,esk3_3(X9,X10,X11))|~r1_orders_2(X9,X11,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(~v3_orders_2(X9)|~l1_orders_2(X9))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t38_orders_2])])])])])])).
cnf(c_0_16, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_21, plain, ![X14, X15]:(v1_xboole_0(X14)|~m1_subset_1(X15,X14)|k6_domain_1(X14,X15)=k1_tarski(X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_22, plain, ![X16]:k2_tarski(X16,X16)=k1_tarski(X16), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_23, plain, ![X17, X18]:k1_enumset1(X17,X17,X18)=k2_tarski(X17,X18), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_24, plain, ![X19, X20, X21]:(~r2_hidden(X19,X20)|(~m1_subset_1(X20,k1_zfmisc_1(X21))|~v1_xboole_0(X21))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
cnf(c_0_25, plain, (r2_hidden(X1,esk3_3(X2,X3,X1))|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~v3_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_26, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18]), c_0_19])]), c_0_20])).
fof(c_0_27, plain, ![X22, X23, X24, X25, X26, X27, X28, X29]:(((~r2_hidden(X25,X24)|(X25=X22|X25=X23)|X24!=k2_tarski(X22,X23))&((X26!=X22|r2_hidden(X26,X24)|X24!=k2_tarski(X22,X23))&(X26!=X23|r2_hidden(X26,X24)|X24!=k2_tarski(X22,X23))))&(((esk4_3(X27,X28,X29)!=X27|~r2_hidden(esk4_3(X27,X28,X29),X29)|X29=k2_tarski(X27,X28))&(esk4_3(X27,X28,X29)!=X28|~r2_hidden(esk4_3(X27,X28,X29),X29)|X29=k2_tarski(X27,X28)))&(r2_hidden(esk4_3(X27,X28,X29),X29)|(esk4_3(X27,X28,X29)=X27|esk4_3(X27,X28,X29)=X28)|X29=k2_tarski(X27,X28)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_28, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_31, plain, ![X31, X32]:(~m1_subset_1(X31,X32)|(v1_xboole_0(X32)|r2_hidden(X31,X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_32, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, negated_conjecture, (r2_hidden(esk2_0,esk3_3(esk1_0,esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_17]), c_0_18]), c_0_19])])).
cnf(c_0_34, plain, (m1_subset_1(esk3_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_35, plain, ![X33, X34, X35]:(v2_struct_0(X33)|(~v3_orders_2(X33)|(~v4_orders_2(X33)|(~v5_orders_2(X33)|~l1_orders_2(X33))))|(~m1_subset_1(X34,u1_struct_0(X33))|(~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X33)))|(~r2_hidden(X34,X35)|~r2_hidden(X34,k1_orders_2(X33,X35)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t47_orders_2])])])])).
fof(c_0_36, plain, ![X36, X37, X38]:(~r2_hidden(X36,X37)|~m1_subset_1(X37,k1_zfmisc_1(X38))|m1_subset_1(X36,X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_37, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_38, plain, (k6_domain_1(X1,X2)=k1_enumset1(X2,X2,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_29]), c_0_30])).
fof(c_0_39, plain, ![X39, X40]:(~r2_hidden(X39,X40)|m1_subset_1(k1_tarski(X39),k1_zfmisc_1(X40))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_subset_1])])).
cnf(c_0_40, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41, negated_conjecture, (~v1_xboole_0(X1)|~m1_subset_1(esk3_3(esk1_0,esk2_0,esk2_0),k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_42, negated_conjecture, (m1_subset_1(esk3_3(esk1_0,esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_26]), c_0_17]), c_0_18]), c_0_19])])).
cnf(c_0_43, plain, (v2_struct_0(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,X3)|~r2_hidden(X2,k1_orders_2(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_44, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_45, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X4,X2)), inference(rw,[status(thm)],[c_0_37, c_0_30])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk2_0,k1_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_47, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k1_enumset1(esk2_0,esk2_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_38, c_0_17])).
cnf(c_0_48, plain, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_49, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_40, c_0_17])).
cnf(c_0_50, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_51, plain, (v2_struct_0(X1)|~r2_hidden(X2,k1_orders_2(X1,X3))|~r2_hidden(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_52, plain, (r2_hidden(X1,k1_enumset1(X2,X2,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_45])])).
cnf(c_0_53, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,k1_orders_2(esk1_0,k1_enumset1(esk2_0,esk2_0,esk2_0)))), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_54, plain, (m1_subset_1(k1_enumset1(X1,X1,X1),k1_zfmisc_1(X2))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_29]), c_0_30])).
cnf(c_0_55, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_56, plain, (v2_struct_0(X1)|~r2_hidden(X2,k1_orders_2(X1,k1_enumset1(X3,X3,X2)))|~m1_subset_1(k1_enumset1(X3,X3,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_57, negated_conjecture, (r2_hidden(esk2_0,k1_orders_2(esk1_0,k1_enumset1(esk2_0,esk2_0,esk2_0)))), inference(sr,[status(thm)],[c_0_53, c_0_50])).
cnf(c_0_58, negated_conjecture, (m1_subset_1(k1_enumset1(esk2_0,esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_59, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_60, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_61, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58]), c_0_18]), c_0_59]), c_0_60]), c_0_19])]), c_0_20]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 62
# Proof object clause steps            : 37
# Proof object formula steps           : 25
# Proof object conjectures             : 22
# Proof object clause conjectures      : 19
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 29
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 29
# Processed clauses                    : 147
# ...of these trivial                  : 8
# ...subsumed                          : 15
# ...remaining for further processing  : 124
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 8
# Generated clauses                    : 126
# ...of the previous two non-trivial   : 115
# Contextual simplify-reflections      : 3
# Paramodulations                      : 117
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 76
#    Positive orientable unit clauses  : 26
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 45
# Current number of unprocessed clauses: 17
# ...number of literals in the above   : 72
# Current number of archived formulas  : 0
# Current number of archived clauses   : 47
# Clause-clause subsumption calls (NU) : 764
# Rec. Clause-clause subsumption calls : 143
# Non-unit clause-clause subsumptions  : 13
# Unit Clause-clause subsumption calls : 135
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5269

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28788 pages
