# No SInE strategy applied
# Auto-Mode selected heuristic G_____0017_C18_F1_SE_CS_SP_S4Y
# and selection function SelectMaxLComplexAPPNTNp.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t138_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=>(k2_zfmisc_1(X1,X2)=k1_xboole_0|(r1_tarski(X1,X3)&r1_tarski(X2,X4)))), file('data/MPTP/proofs/474d65c6ec6f4b86a8d5e859bd6744cb.E_input', t138_zfmisc_1)).
fof(t134_zfmisc_1, axiom, ![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4)))), file('data/MPTP/proofs/474d65c6ec6f4b86a8d5e859bd6744cb.E_input', t134_zfmisc_1)).
fof(t123_zfmisc_1, axiom, ![X1, X2, X3, X4]:k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)), file('data/MPTP/proofs/474d65c6ec6f4b86a8d5e859bd6744cb.E_input', t123_zfmisc_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/474d65c6ec6f4b86a8d5e859bd6744cb.E_input', t28_xboole_1)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/474d65c6ec6f4b86a8d5e859bd6744cb.E_input', t17_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/474d65c6ec6f4b86a8d5e859bd6744cb.E_input', commutativity_k3_xboole_0)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/474d65c6ec6f4b86a8d5e859bd6744cb.E_input', t113_zfmisc_1)).
fof(c_0_7, negated_conjecture, ~(![X1, X2, X3, X4]:(r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=>(k2_zfmisc_1(X1,X2)=k1_xboole_0|(r1_tarski(X1,X3)&r1_tarski(X2,X4))))), inference(assume_negation,[status(cth)],[t138_zfmisc_1])).
fof(c_0_8, plain, ![X9, X10, X11, X12]:((X9=X11|X10=k1_xboole_0|X9=k1_xboole_0|k2_zfmisc_1(X9,X10)!=k2_zfmisc_1(X11,X12))&(X10=X12|X10=k1_xboole_0|X9=k1_xboole_0|k2_zfmisc_1(X9,X10)!=k2_zfmisc_1(X11,X12))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t134_zfmisc_1])])])).
fof(c_0_9, plain, ![X13, X14, X15, X16]:k2_zfmisc_1(k3_xboole_0(X13,X14),k3_xboole_0(X15,X16))=k3_xboole_0(k2_zfmisc_1(X13,X15),k2_zfmisc_1(X14,X16)), inference(variable_rename,[status(thm)],[t123_zfmisc_1])).
fof(c_0_10, plain, ![X17, X18]:(~r1_tarski(X17,X18)|k3_xboole_0(X17,X18)=X17), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_11, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))&(k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0&(~r1_tarski(esk1_0,esk3_0)|~r1_tarski(esk2_0,esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
cnf(c_0_12, plain, (X1=X2|X1=k1_xboole_0|X3=k1_xboole_0|k2_zfmisc_1(X3,X1)!=k2_zfmisc_1(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_16, plain, ![X19, X20]:r1_tarski(k3_xboole_0(X19,X20),X19), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_17, plain, ![X21, X22]:k3_xboole_0(X21,X22)=k3_xboole_0(X22,X21), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_18, plain, (X1=k3_xboole_0(X2,X3)|X4=k1_xboole_0|X1=k1_xboole_0|k2_zfmisc_1(X4,X1)!=k3_xboole_0(k2_zfmisc_1(X5,X2),k2_zfmisc_1(X6,X3))), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_19, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))=k2_zfmisc_1(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_20, plain, (X1=X2|X3=k1_xboole_0|X1=k1_xboole_0|k2_zfmisc_1(X1,X3)!=k2_zfmisc_1(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_21, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, negated_conjecture, (X1=k3_xboole_0(esk2_0,esk4_0)|X1=k1_xboole_0|X2=k1_xboole_0|k2_zfmisc_1(X2,X1)!=k2_zfmisc_1(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
fof(c_0_24, plain, ![X23, X24]:((k2_zfmisc_1(X23,X24)!=k1_xboole_0|(X23=k1_xboole_0|X24=k1_xboole_0))&((X23!=k1_xboole_0|k2_zfmisc_1(X23,X24)=k1_xboole_0)&(X24!=k1_xboole_0|k2_zfmisc_1(X23,X24)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_25, plain, (X1=k3_xboole_0(X2,X3)|X4=k1_xboole_0|X1=k1_xboole_0|k2_zfmisc_1(X1,X4)!=k3_xboole_0(k2_zfmisc_1(X2,X5),k2_zfmisc_1(X3,X6))), inference(spm,[status(thm)],[c_0_20, c_0_13])).
cnf(c_0_26, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27, negated_conjecture, (k3_xboole_0(esk2_0,esk4_0)=esk2_0|esk1_0=k1_xboole_0|esk2_0=k1_xboole_0), inference(er,[status(thm)],[c_0_23])).
cnf(c_0_28, plain, (X1=k1_xboole_0|X2=k1_xboole_0|k2_zfmisc_1(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29, negated_conjecture, (X1=k3_xboole_0(esk1_0,esk3_0)|X1=k1_xboole_0|X2=k1_xboole_0|k2_zfmisc_1(X1,X2)!=k2_zfmisc_1(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_25, c_0_19])).
cnf(c_0_30, negated_conjecture, (~r1_tarski(esk1_0,esk3_0)|~r1_tarski(esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_31, negated_conjecture, (esk2_0=k1_xboole_0|esk1_0=k1_xboole_0|r1_tarski(esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|k3_xboole_0(X3,X4)=k1_xboole_0|k3_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X4,X2))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_28, c_0_13])).
cnf(c_0_33, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),X1),k2_zfmisc_1(k2_zfmisc_1(esk3_0,esk4_0),X2))=k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_13, c_0_19])).
cnf(c_0_34, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_35, negated_conjecture, (k3_xboole_0(esk1_0,esk3_0)=esk1_0|esk2_0=k1_xboole_0|esk1_0=k1_xboole_0), inference(er,[status(thm)],[c_0_29])).
cnf(c_0_36, negated_conjecture, (esk1_0=k1_xboole_0|esk2_0=k1_xboole_0|~r1_tarski(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_37, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_38, negated_conjecture, (k3_xboole_0(X1,X2)=k1_xboole_0|k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),k3_xboole_0(X1,X2))!=k1_xboole_0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_19]), c_0_34])).
cnf(c_0_39, negated_conjecture, (esk1_0=k1_xboole_0|esk2_0=k1_xboole_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_35]), c_0_36])).
cnf(c_0_40, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_41, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_42, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk1_0,esk2_0))!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_19]), c_0_34])).
cnf(c_0_43, negated_conjecture, (esk1_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_39]), c_0_40])])).
cnf(c_0_44, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_45, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_43]), c_0_44]), c_0_43]), c_0_44]), c_0_44])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 46
# Proof object clause steps            : 31
# Proof object formula steps           : 15
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 14
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 12
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 12
# Processed clauses                    : 28938
# ...of these trivial                  : 3846
# ...subsumed                          : 23782
# ...remaining for further processing  : 1310
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 230
# Backward-rewritten                   : 569
# Generated clauses                    : 589334
# ...of the previous two non-trivial   : 290348
# Contextual simplify-reflections      : 3
# Paramodulations                      : 589236
# Factorizations                       : 82
# Equation resolutions                 : 16
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 509
#    Positive orientable unit clauses  : 394
#    Positive unorientable unit clauses: 20
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 95
# Current number of unprocessed clauses: 255884
# ...number of literals in the above   : 448892
# Current number of archived formulas  : 0
# Current number of archived clauses   : 799
# Clause-clause subsumption calls (NU) : 200382
# Rec. Clause-clause subsumption calls : 186430
# Non-unit clause-clause subsumptions  : 23713
# Unit Clause-clause subsumption calls : 336
# Rewrite failures with RHS unbound    : 10
# BW rewrite match attempts            : 14843
# BW rewrite match successes           : 530
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 9835347

# -------------------------------------------------
# User time                : 3.406 s
# System time              : 0.076 s
# Total time               : 3.482 s
# Maximum resident set size: 195088 pages
