# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/479ac45112cc42d08bb8b09e4f94c8c1.E_input', d10_xboole_0)).
fof(t1_lattice3, axiom, ![X1, X2]:(m1_subset_1(X2,u1_struct_0(k1_lattice3(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k1_lattice3(X1)))=>(k1_lattices(k1_lattice3(X1),X2,X3)=k2_xboole_0(X2,X3)&k2_lattices(k1_lattice3(X1),X2,X3)=k3_xboole_0(X2,X3)))), file('data/MPTP/proofs/479ac45112cc42d08bb8b09e4f94c8c1.E_input', t1_lattice3)).
fof(d3_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l2_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattices(X1,X2,X3)<=>k1_lattices(X1,X2,X3)=X3)))), file('data/MPTP/proofs/479ac45112cc42d08bb8b09e4f94c8c1.E_input', d3_lattices)).
fof(fc1_lattice3, axiom, ![X1]:(~(v2_struct_0(k1_lattice3(X1)))&v3_lattices(k1_lattice3(X1))), file('data/MPTP/proofs/479ac45112cc42d08bb8b09e4f94c8c1.E_input', fc1_lattice3)).
fof(t2_lattice3, conjecture, ![X1, X2]:(m1_subset_1(X2,u1_struct_0(k1_lattice3(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k1_lattice3(X1)))=>(r1_lattices(k1_lattice3(X1),X2,X3)<=>r1_tarski(X2,X3)))), file('data/MPTP/proofs/479ac45112cc42d08bb8b09e4f94c8c1.E_input', t2_lattice3)).
fof(t11_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k2_xboole_0(X1,X2),X3)=>r1_tarski(X1,X3)), file('data/MPTP/proofs/479ac45112cc42d08bb8b09e4f94c8c1.E_input', t11_xboole_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/479ac45112cc42d08bb8b09e4f94c8c1.E_input', t12_xboole_1)).
fof(dt_l3_lattices, axiom, ![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1))), file('data/MPTP/proofs/479ac45112cc42d08bb8b09e4f94c8c1.E_input', dt_l3_lattices)).
fof(dt_k1_lattice3, axiom, ![X1]:(v3_lattices(k1_lattice3(X1))&l3_lattices(k1_lattice3(X1))), file('data/MPTP/proofs/479ac45112cc42d08bb8b09e4f94c8c1.E_input', dt_k1_lattice3)).
fof(c_0_9, plain, ![X17, X18]:(((r1_tarski(X17,X18)|X17!=X18)&(r1_tarski(X18,X17)|X17!=X18))&(~r1_tarski(X17,X18)|~r1_tarski(X18,X17)|X17=X18)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_10, plain, ![X20, X21, X22]:((k1_lattices(k1_lattice3(X20),X21,X22)=k2_xboole_0(X21,X22)|~m1_subset_1(X22,u1_struct_0(k1_lattice3(X20)))|~m1_subset_1(X21,u1_struct_0(k1_lattice3(X20))))&(k2_lattices(k1_lattice3(X20),X21,X22)=k3_xboole_0(X21,X22)|~m1_subset_1(X22,u1_struct_0(k1_lattice3(X20)))|~m1_subset_1(X21,u1_struct_0(k1_lattice3(X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_lattice3])])])])).
fof(c_0_11, plain, ![X9, X10, X11]:((~r1_lattices(X9,X10,X11)|k1_lattices(X9,X10,X11)=X11|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(v2_struct_0(X9)|~l2_lattices(X9)))&(k1_lattices(X9,X10,X11)!=X11|r1_lattices(X9,X10,X11)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(v2_struct_0(X9)|~l2_lattices(X9)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d3_lattices])])])])])).
fof(c_0_12, plain, ![X13]:(~v2_struct_0(k1_lattice3(X13))&v3_lattices(k1_lattice3(X13))), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_lattice3])])).
fof(c_0_13, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,u1_struct_0(k1_lattice3(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k1_lattice3(X1)))=>(r1_lattices(k1_lattice3(X1),X2,X3)<=>r1_tarski(X2,X3))))), inference(assume_negation,[status(cth)],[t2_lattice3])).
fof(c_0_14, plain, ![X14, X15, X16]:(~r1_tarski(k2_xboole_0(X14,X15),X16)|r1_tarski(X14,X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_xboole_1])])).
cnf(c_0_15, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, plain, (k1_lattices(k1_lattice3(X1),X2,X3)=k2_xboole_0(X2,X3)|~m1_subset_1(X3,u1_struct_0(k1_lattice3(X1)))|~m1_subset_1(X2,u1_struct_0(k1_lattice3(X1)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, plain, (k1_lattices(X1,X2,X3)=X3|v2_struct_0(X1)|~r1_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, plain, (~v2_struct_0(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_19, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(k1_lattice3(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(k1_lattice3(esk1_0)))&((~r1_lattices(k1_lattice3(esk1_0),esk2_0,esk3_0)|~r1_tarski(esk2_0,esk3_0))&(r1_lattices(k1_lattice3(esk1_0),esk2_0,esk3_0)|r1_tarski(esk2_0,esk3_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
fof(c_0_20, plain, ![X7, X8]:(~r1_tarski(X7,X8)|k2_xboole_0(X7,X8)=X8), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_21, plain, (r1_tarski(X1,X3)|~r1_tarski(k2_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (k2_xboole_0(X1,X2)=X2|~l2_lattices(k1_lattice3(X3))|~r1_lattices(k1_lattice3(X3),X1,X2)|~m1_subset_1(X2,u1_struct_0(k1_lattice3(X3)))|~m1_subset_1(X1,u1_struct_0(k1_lattice3(X3)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])).
cnf(c_0_24, negated_conjecture, (r1_lattices(k1_lattice3(esk1_0),esk2_0,esk3_0)|r1_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k1_lattice3(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(k1_lattice3(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_29, negated_conjecture, (k2_xboole_0(esk2_0,esk3_0)=esk3_0|~l2_lattices(k1_lattice3(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
fof(c_0_30, plain, ![X19]:((l1_lattices(X19)|~l3_lattices(X19))&(l2_lattices(X19)|~l3_lattices(X19))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
fof(c_0_31, plain, ![X12]:(v3_lattices(k1_lattice3(X12))&l3_lattices(k1_lattice3(X12))), inference(variable_rename,[status(thm)],[dt_k1_lattice3])).
cnf(c_0_32, negated_conjecture, (r1_tarski(esk2_0,esk3_0)|~l2_lattices(k1_lattice3(esk1_0))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_33, plain, (l2_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_34, plain, (l3_lattices(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_35, negated_conjecture, (~r1_lattices(k1_lattice3(esk1_0),esk2_0,esk3_0)|~r1_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_36, negated_conjecture, (r1_tarski(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])])).
cnf(c_0_37, plain, (r1_lattices(X1,X2,X3)|v2_struct_0(X1)|k1_lattices(X1,X2,X3)!=X3|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_38, negated_conjecture, (~r1_lattices(k1_lattice3(esk1_0),esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_36])])).
cnf(c_0_39, plain, (r1_lattices(k1_lattice3(X1),X2,X3)|k2_xboole_0(X2,X3)!=X3|~l2_lattices(k1_lattice3(X1))|~m1_subset_1(X3,u1_struct_0(k1_lattice3(X1)))|~m1_subset_1(X2,u1_struct_0(k1_lattice3(X1)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_16]), c_0_18])).
cnf(c_0_40, negated_conjecture, (~l2_lattices(k1_lattice3(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_25]), c_0_26])]), c_0_29])).
cnf(c_0_41, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_33]), c_0_34])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 42
# Proof object clause steps            : 23
# Proof object formula steps           : 19
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 19
# Processed clauses                    : 129
# ...of these trivial                  : 2
# ...subsumed                          : 46
# ...remaining for further processing  : 81
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 3
# Generated clauses                    : 212
# ...of the previous two non-trivial   : 187
# Contextual simplify-reflections      : 2
# Paramodulations                      : 210
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 57
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 42
# Current number of unprocessed clauses: 92
# ...number of literals in the above   : 208
# Current number of archived formulas  : 0
# Current number of archived clauses   : 22
# Clause-clause subsumption calls (NU) : 857
# Rec. Clause-clause subsumption calls : 711
# Non-unit clause-clause subsumptions  : 49
# Unit Clause-clause subsumption calls : 13
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3541

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28684 pages
