# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t28_zfmisc_1, conjecture, ![X1, X2, X3, X4]:~((r1_tarski(k2_tarski(X1,X2),k2_tarski(X3,X4))&(~(X1=X3)&~(X1=X4)))), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t28_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t75_enumset1)).
fof(t18_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k3_xboole_0(X2,X3))=>r1_tarski(X1,X2)), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t18_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', commutativity_k3_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t28_xboole_1)).
fof(t12_zfmisc_1, axiom, ![X1, X2]:r1_tarski(k1_tarski(X1),k2_tarski(X1,X2)), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t12_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t69_enumset1)).
fof(t68_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k5_enumset1(X1,X2,X3,X4,X5,X6,X7),k1_tarski(X8)), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t68_enumset1)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t98_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t100_xboole_1)).
fof(t134_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8),k1_tarski(X9)), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t134_enumset1)).
fof(t107_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X4,X3,X2), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t107_enumset1)).
fof(t92_xboole_1, axiom, ![X1]:k5_xboole_0(X1,X1)=k1_xboole_0, file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t92_xboole_1)).
fof(t5_boole, axiom, ![X1]:k5_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t5_boole)).
fof(t103_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X2,X4,X3), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t103_enumset1)).
fof(t55_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_tarski(X6)), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t55_enumset1)).
fof(t66_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_enumset1(X6,X7,X8)), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t66_enumset1)).
fof(t132_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k4_enumset1(X1,X2,X3,X4,X5,X6),k1_enumset1(X7,X8,X9)), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t132_enumset1)).
fof(t131_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k2_enumset1(X6,X7,X8,X9)), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t131_enumset1)).
fof(t125_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t125_enumset1)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t17_xboole_1)).
fof(t69_xboole_1, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>~((r1_tarski(X2,X1)&r1_xboole_0(X2,X1)))), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t69_xboole_1)).
fof(l97_xboole_1, axiom, ![X1, X2]:r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2)), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', l97_xboole_1)).
fof(fc3_xboole_0, axiom, ![X1, X2]:~(v1_xboole_0(k2_tarski(X1,X2))), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', fc3_xboole_0)).
fof(commutativity_k5_xboole_0, axiom, ![X1, X2]:k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', commutativity_k5_xboole_0)).
fof(t91_xboole_1, axiom, ![X1, X2, X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3)), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t91_xboole_1)).
fof(t136_enumset1, axiom, ![X1, X2, X3]:~((~(X1=X2)&(~(X1=X3)&~(k4_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X1))=k2_tarski(X2,X3))))), file('data/MPTP/proofs/48ab212ffb6f426ba9a1d36277d28361.E_input', t136_enumset1)).
fof(c_0_32, negated_conjecture, ~(![X1, X2, X3, X4]:~((r1_tarski(k2_tarski(X1,X2),k2_tarski(X3,X4))&(~(X1=X3)&~(X1=X4))))), inference(assume_negation,[status(cth)],[t28_zfmisc_1])).
fof(c_0_33, negated_conjecture, (r1_tarski(k2_tarski(esk1_0,esk2_0),k2_tarski(esk3_0,esk4_0))&(esk1_0!=esk3_0&esk1_0!=esk4_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_32])])])])).
fof(c_0_34, plain, ![X68, X69]:k1_enumset1(X68,X68,X69)=k2_tarski(X68,X69), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_35, plain, ![X97, X98, X99]:k2_enumset1(X97,X97,X98,X99)=k1_enumset1(X97,X98,X99), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_36, plain, ![X48, X49, X50, X51]:k3_enumset1(X48,X48,X49,X50,X51)=k2_enumset1(X48,X49,X50,X51), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_37, plain, ![X70, X71, X72, X73, X74]:k4_enumset1(X70,X70,X71,X72,X73,X74)=k3_enumset1(X70,X71,X72,X73,X74), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_38, plain, ![X76, X77, X78, X79, X80, X81]:k5_enumset1(X76,X76,X77,X78,X79,X80,X81)=k4_enumset1(X76,X77,X78,X79,X80,X81), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_39, plain, ![X112, X113, X114, X115, X116, X117, X118]:k6_enumset1(X112,X112,X113,X114,X115,X116,X117,X118)=k5_enumset1(X112,X113,X114,X115,X116,X117,X118), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_40, plain, ![X100, X101, X102]:(~r1_tarski(X100,k3_xboole_0(X101,X102))|r1_tarski(X100,X101)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t18_xboole_1])])).
fof(c_0_41, plain, ![X128, X129]:k3_xboole_0(X128,X129)=k3_xboole_0(X129,X128), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_42, plain, ![X32, X33]:(~r1_tarski(X32,X33)|k3_xboole_0(X32,X33)=X32), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_43, negated_conjecture, (r1_tarski(k2_tarski(esk1_0,esk2_0),k2_tarski(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_44, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_46, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_47, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_48, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_49, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_50, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_51, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_52, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_53, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0),k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49])).
cnf(c_0_54, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k3_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_55, negated_conjecture, (k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0),k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
fof(c_0_56, plain, ![X126, X127]:r1_tarski(k1_tarski(X126),k2_tarski(X126,X127)), inference(variable_rename,[status(thm)],[t12_zfmisc_1])).
fof(c_0_57, plain, ![X94]:k2_tarski(X94,X94)=k1_tarski(X94), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_58, plain, ![X40, X41, X42, X43, X44, X45, X46, X47]:k6_enumset1(X40,X41,X42,X43,X44,X45,X46,X47)=k2_xboole_0(k5_enumset1(X40,X41,X42,X43,X44,X45,X46),k1_tarski(X47)), inference(variable_rename,[status(thm)],[t68_enumset1])).
fof(c_0_59, plain, ![X55, X56]:k2_xboole_0(X55,X56)=k5_xboole_0(X55,k4_xboole_0(X56,X55)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_60, plain, ![X124, X125]:k4_xboole_0(X124,X125)=k5_xboole_0(X124,k3_xboole_0(X124,X125)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_61, plain, ![X23, X24, X25, X26, X27, X28, X29, X30, X31]:k7_enumset1(X23,X24,X25,X26,X27,X28,X29,X30,X31)=k2_xboole_0(k6_enumset1(X23,X24,X25,X26,X27,X28,X29,X30),k1_tarski(X31)), inference(variable_rename,[status(thm)],[t134_enumset1])).
cnf(c_0_62, negated_conjecture, (r1_tarski(X1,k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))|~r1_tarski(X1,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_63, plain, (r1_tarski(k1_tarski(X1),k2_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_64, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_65, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k5_enumset1(X1,X2,X3,X4,X5,X6,X7),k1_tarski(X8))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_66, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_67, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_68, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8),k1_tarski(X9))), inference(split_conjunct,[status(thm)],[c_0_61])).
fof(c_0_69, plain, ![X106, X107, X108, X109]:k2_enumset1(X106,X107,X108,X109)=k2_enumset1(X106,X109,X108,X107), inference(variable_rename,[status(thm)],[t107_enumset1])).
cnf(c_0_70, negated_conjecture, (k3_xboole_0(X1,k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))=X1|~r1_tarski(X1,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_52, c_0_62])).
cnf(c_0_71, plain, (r1_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_64]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49])).
fof(c_0_72, plain, ![X75]:k5_xboole_0(X75,X75)=k1_xboole_0, inference(variable_rename,[status(thm)],[t92_xboole_1])).
fof(c_0_73, plain, ![X123]:k5_xboole_0(X123,k1_xboole_0)=X123, inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_74, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k5_xboole_0(k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7),k5_xboole_0(k6_enumset1(X8,X8,X8,X8,X8,X8,X8,X8),k3_xboole_0(k6_enumset1(X8,X8,X8,X8,X8,X8,X8,X8),k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_64]), c_0_44]), c_0_66]), c_0_67]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_49])).
cnf(c_0_75, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k5_xboole_0(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8),k5_xboole_0(k6_enumset1(X9,X9,X9,X9,X9,X9,X9,X9),k3_xboole_0(k6_enumset1(X9,X9,X9,X9,X9,X9,X9,X9),k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_64]), c_0_44]), c_0_66]), c_0_67]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49])).
cnf(c_0_76, plain, (k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X4,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
fof(c_0_77, plain, ![X130, X131, X132, X133]:k2_enumset1(X130,X131,X132,X133)=k2_enumset1(X130,X131,X133,X132), inference(variable_rename,[status(thm)],[t103_enumset1])).
fof(c_0_78, plain, ![X34, X35, X36, X37, X38, X39]:k4_enumset1(X34,X35,X36,X37,X38,X39)=k2_xboole_0(k3_enumset1(X34,X35,X36,X37,X38),k1_tarski(X39)), inference(variable_rename,[status(thm)],[t55_enumset1])).
fof(c_0_79, plain, ![X82, X83, X84, X85, X86, X87, X88, X89]:k6_enumset1(X82,X83,X84,X85,X86,X87,X88,X89)=k2_xboole_0(k3_enumset1(X82,X83,X84,X85,X86),k1_enumset1(X87,X88,X89)), inference(variable_rename,[status(thm)],[t66_enumset1])).
cnf(c_0_80, negated_conjecture, (k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_81, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_82, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_83, plain, (k7_enumset1(X1,X1,X2,X3,X4,X5,X6,X7,X8)=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)), inference(rw,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_84, plain, (k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k6_enumset1(X1,X1,X1,X1,X1,X4,X3,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49])).
cnf(c_0_85, plain, (k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X2,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_86, plain, (k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_tarski(X6))), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_87, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_enumset1(X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_88, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)=k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_80]), c_0_81]), c_0_82]), c_0_83]), c_0_84])).
cnf(c_0_89, plain, (k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k6_enumset1(X1,X1,X1,X1,X1,X2,X4,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49])).
cnf(c_0_90, plain, (k6_enumset1(X1,X1,X1,X2,X3,X4,X5,X6)=k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5),k5_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X6,X6,X6),k3_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X6,X6,X6),k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_64]), c_0_44]), c_0_66]), c_0_67]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_48]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_49]), c_0_49])).
cnf(c_0_91, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5),k5_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X6,X7,X8),k3_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X6,X7,X8),k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_66]), c_0_67]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_49])).
fof(c_0_92, plain, ![X57, X58, X59, X60, X61, X62, X63, X64, X65]:k7_enumset1(X57,X58,X59,X60,X61,X62,X63,X64,X65)=k2_xboole_0(k4_enumset1(X57,X58,X59,X60,X61,X62),k1_enumset1(X63,X64,X65)), inference(variable_rename,[status(thm)],[t132_enumset1])).
cnf(c_0_93, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)=k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk1_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_94, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X6,X6)=k6_enumset1(X1,X1,X1,X2,X3,X4,X5,X6)), inference(rw,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_95, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k4_enumset1(X1,X2,X3,X4,X5,X6),k1_enumset1(X7,X8,X9))), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_96, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_94]), c_0_94])).
cnf(c_0_97, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k5_xboole_0(k6_enumset1(X1,X1,X1,X2,X3,X4,X5,X6),k5_xboole_0(k6_enumset1(X7,X7,X7,X7,X7,X7,X8,X9),k3_xboole_0(k6_enumset1(X7,X7,X7,X7,X7,X7,X8,X9),k6_enumset1(X1,X1,X1,X2,X3,X4,X5,X6))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_95, c_0_66]), c_0_67]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_49])).
cnf(c_0_98, negated_conjecture, (k7_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,X1)=k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_96]), c_0_75]), c_0_83])).
cnf(c_0_99, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X7,X7)=k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_97]), c_0_83])).
cnf(c_0_100, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,esk4_0)=k6_enumset1(esk3_0,esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0)), inference(spm,[status(thm)],[c_0_98, c_0_99])).
fof(c_0_101, plain, ![X14, X15, X16, X17, X18, X19, X20, X21, X22]:k7_enumset1(X14,X15,X16,X17,X18,X19,X20,X21,X22)=k2_xboole_0(k3_enumset1(X14,X15,X16,X17,X18),k2_enumset1(X19,X20,X21,X22)), inference(variable_rename,[status(thm)],[t131_enumset1])).
cnf(c_0_102, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,esk4_0,X1)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_100]), c_0_75]), c_0_83]), c_0_83])).
cnf(c_0_103, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k2_enumset1(X6,X7,X8,X9))), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_104, plain, (k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5)=k6_enumset1(X1,X1,X1,X1,X4,X3,X2,X5)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_84]), c_0_75]), c_0_83]), c_0_83])).
cnf(c_0_105, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,X1,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,X1)), inference(spm,[status(thm)],[c_0_89, c_0_102])).
cnf(c_0_106, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5),k5_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X7,X8,X9),k3_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X7,X8,X9),k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_66]), c_0_67]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_49])).
cnf(c_0_107, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,X1,esk4_0,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,X1)), inference(spm,[status(thm)],[c_0_84, c_0_102])).
cnf(c_0_108, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,X1,esk4_0,esk3_0,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,X1)), inference(spm,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_109, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X5,X5)=k6_enumset1(X1,X1,X1,X1,X2,X4,X3,X5)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_89]), c_0_97]), c_0_83]), c_0_83])).
cnf(c_0_110, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X6,X7,X8)=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_106]), c_0_83])).
cnf(c_0_111, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,esk3_0)), inference(rw,[status(thm)],[c_0_100, c_0_107])).
cnf(c_0_112, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,X1,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,X1)), inference(spm,[status(thm)],[c_0_108, c_0_109])).
cnf(c_0_113, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,X1)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,X1)), inference(rw,[status(thm)],[c_0_98, c_0_110])).
cnf(c_0_114, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,esk3_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,esk1_0)), inference(rw,[status(thm)],[c_0_111, c_0_112])).
cnf(c_0_115, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,esk1_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_113]), c_0_114])).
fof(c_0_116, plain, ![X119, X120, X121, X122]:k2_enumset1(X119,X120,X121,X122)=k2_enumset1(X122,X121,X120,X119), inference(variable_rename,[status(thm)],[t125_enumset1])).
cnf(c_0_117, negated_conjecture, (k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_80, c_0_88]), c_0_89]), c_0_94]), c_0_94])).
cnf(c_0_118, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,X1)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk1_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_115]), c_0_75]), c_0_110]), c_0_110])).
cnf(c_0_119, plain, (k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_116])).
cnf(c_0_120, negated_conjecture, (k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk1_0,esk1_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_117, c_0_115]), c_0_118])).
cnf(c_0_121, plain, (k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k6_enumset1(X4,X4,X4,X4,X4,X3,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_119, c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49])).
cnf(c_0_122, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk1_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_120]), c_0_81]), c_0_82]), c_0_99])).
cnf(c_0_123, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k6_enumset1(X4,X3,X2,X1,X5,X6,X7,X8)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_121]), c_0_106]), c_0_83]), c_0_83])).
cnf(c_0_124, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk1_0,esk1_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk4_0,esk4_0,esk4_0,esk1_0)), inference(rw,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_125, negated_conjecture, (k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk4_0,esk4_0,esk4_0,esk1_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[c_0_120, c_0_124])).
cnf(c_0_126, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk4_0,esk4_0,esk4_0,esk1_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk4_0,esk4_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_125]), c_0_81]), c_0_82]), c_0_110])).
cnf(c_0_127, negated_conjecture, (k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk4_0,esk4_0,esk1_0,esk1_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[c_0_125, c_0_126])).
cnf(c_0_128, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk4_0,esk4_0,esk1_0,esk1_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk4_0,esk4_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_127]), c_0_81]), c_0_82]), c_0_99]), c_0_123])).
cnf(c_0_129, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk4_0,esk4_0,esk4_0,esk1_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk4_0,esk4_0,esk1_0)), inference(rw,[status(thm)],[c_0_126, c_0_128])).
cnf(c_0_130, negated_conjecture, (r1_tarski(k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0),k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk1_0,esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_71, c_0_93])).
fof(c_0_131, plain, ![X66, X67]:r1_tarski(k3_xboole_0(X66,X67),X66), inference(variable_rename,[status(thm)],[t17_xboole_1])).
cnf(c_0_132, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk4_0,esk4_0,esk1_0,X1)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_129]), c_0_75]), c_0_110]), c_0_110])).
cnf(c_0_133, negated_conjecture, (r1_tarski(k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_130, c_0_94]), c_0_94])).
cnf(c_0_134, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_131])).
fof(c_0_135, plain, ![X110, X111]:(v1_xboole_0(X111)|(~r1_tarski(X111,X110)|~r1_xboole_0(X111,X110))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t69_xboole_1])])])).
fof(c_0_136, plain, ![X95, X96]:r1_xboole_0(k3_xboole_0(X95,X96),k5_xboole_0(X95,X96)), inference(variable_rename,[status(thm)],[l97_xboole_1])).
fof(c_0_137, plain, ![X90, X91]:~v1_xboole_0(k2_tarski(X90,X91)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc3_xboole_0])])).
cnf(c_0_138, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk4_0,esk4_0,esk1_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[c_0_128, c_0_132])).
cnf(c_0_139, negated_conjecture, (k3_xboole_0(k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0),k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))=k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_133]), c_0_51])).
cnf(c_0_140, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_134, c_0_51])).
cnf(c_0_141, plain, (v1_xboole_0(X1)|~r1_tarski(X1,X2)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_135])).
cnf(c_0_142, plain, (r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_136])).
cnf(c_0_143, plain, (~v1_xboole_0(k2_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_137])).
cnf(c_0_144, negated_conjecture, (k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0,esk1_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_127, c_0_128]), c_0_138])).
cnf(c_0_145, negated_conjecture, (k3_xboole_0(k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))=k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_139, c_0_115]), c_0_118])).
cnf(c_0_146, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_140, c_0_80])).
cnf(c_0_147, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~r1_tarski(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_141, c_0_142])).
cnf(c_0_148, plain, (~v1_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_143, c_0_44]), c_0_45]), c_0_46]), c_0_47]), c_0_48]), c_0_49])).
cnf(c_0_149, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0,esk1_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_144]), c_0_81]), c_0_82]), c_0_99]), c_0_123])).
cnf(c_0_150, negated_conjecture, (k3_xboole_0(k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk4_0,esk4_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))=k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_145, c_0_124]), c_0_126])).
cnf(c_0_151, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk1_0,esk3_0,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_146, c_0_88]), c_0_89])).
cnf(c_0_152, negated_conjecture, (~r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_117]), c_0_148])).
cnf(c_0_153, negated_conjecture, (k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[c_0_144, c_0_149])).
cnf(c_0_154, negated_conjecture, (k3_xboole_0(k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))=k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_150, c_0_128]), c_0_138])).
cnf(c_0_155, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_151, c_0_94]), c_0_94])).
cnf(c_0_156, negated_conjecture, (~r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,esk1_0)))), inference(rw,[status(thm)],[c_0_152, c_0_115])).
cnf(c_0_157, negated_conjecture, (k7_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0,esk1_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_153]), c_0_81]), c_0_82])).
cnf(c_0_158, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k7_enumset1(X1,X2,X3,X4,X5,X9,X8,X7,X6)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_121]), c_0_106])).
cnf(c_0_159, plain, (k5_xboole_0(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8),k5_xboole_0(k6_enumset1(X9,X9,X9,X9,X9,X9,X9,X9),k3_xboole_0(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8),k6_enumset1(X9,X9,X9,X9,X9,X9,X9,X9))))=k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)), inference(spm,[status(thm)],[c_0_75, c_0_51])).
cnf(c_0_160, negated_conjecture, (k3_xboole_0(k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0),k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))=k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)), inference(rw,[status(thm)],[c_0_154, c_0_149])).
cnf(c_0_161, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk4_0,esk1_0))), inference(rw,[status(thm)],[c_0_155, c_0_115])).
cnf(c_0_162, negated_conjecture, (~r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk1_0,esk1_0)))), inference(rw,[status(thm)],[c_0_156, c_0_118])).
cnf(c_0_163, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_157, c_0_158]), c_0_110])).
cnf(c_0_164, negated_conjecture, (k7_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0,esk1_0,esk3_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_154]), c_0_81]), c_0_82])).
cnf(c_0_165, negated_conjecture, (k7_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0,esk3_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_160]), c_0_81]), c_0_82])).
cnf(c_0_166, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X9,X8)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_89]), c_0_97])).
cnf(c_0_167, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk1_0,esk1_0))), inference(rw,[status(thm)],[c_0_161, c_0_118])).
cnf(c_0_168, negated_conjecture, (~r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk4_0,esk4_0,esk4_0,esk1_0)))), inference(rw,[status(thm)],[c_0_162, c_0_124])).
cnf(c_0_169, negated_conjecture, (k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk3_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[c_0_153, c_0_163])).
cnf(c_0_170, negated_conjecture, (k7_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0,esk1_0,esk3_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0)), inference(rw,[status(thm)],[c_0_164, c_0_149])).
cnf(c_0_171, negated_conjecture, (k7_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0,esk1_0,esk3_0,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_165, c_0_158]), c_0_166]), c_0_163])).
cnf(c_0_172, plain, (k6_enumset1(X1,X2,X3,X4,X5,X5,X6,X7)=k7_enumset1(X1,X2,X3,X4,X4,X4,X5,X6,X7)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_94]), c_0_106]), c_0_83])).
cnf(c_0_173, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk4_0,esk4_0,esk4_0,esk1_0))), inference(rw,[status(thm)],[c_0_167, c_0_124])).
cnf(c_0_174, negated_conjecture, (~r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk4_0,esk4_0,esk1_0,esk1_0)))), inference(rw,[status(thm)],[c_0_168, c_0_126])).
cnf(c_0_175, negated_conjecture, (k3_xboole_0(k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk3_0),k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))=k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)), inference(rw,[status(thm)],[c_0_160, c_0_163])).
cnf(c_0_176, negated_conjecture, (k7_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk1_0,esk3_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_169]), c_0_81]), c_0_82]), c_0_166])).
cnf(c_0_177, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k7_enumset1(X1,X2,X3,X4,X5,X6,X9,X8,X7)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_84]), c_0_106])).
cnf(c_0_178, negated_conjecture, (k7_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0,esk1_0,esk1_0,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_170, c_0_158]), c_0_163])).
cnf(c_0_179, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk3_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk1_0,esk1_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_171, c_0_172])).
fof(c_0_180, plain, ![X92, X93]:k5_xboole_0(X92,X93)=k5_xboole_0(X93,X92), inference(variable_rename,[status(thm)],[commutativity_k5_xboole_0])).
fof(c_0_181, plain, ![X52, X53, X54]:k5_xboole_0(k5_xboole_0(X52,X53),X54)=k5_xboole_0(X52,k5_xboole_0(X53,X54)), inference(variable_rename,[status(thm)],[t91_xboole_1])).
cnf(c_0_182, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk4_0,esk4_0,esk1_0,esk1_0))), inference(rw,[status(thm)],[c_0_173, c_0_126])).
cnf(c_0_183, negated_conjecture, (~r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0,esk1_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_174, c_0_128]), c_0_138])).
cnf(c_0_184, negated_conjecture, (k7_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk3_0,esk3_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_175]), c_0_81]), c_0_82])).
cnf(c_0_185, negated_conjecture, (k7_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk3_0,esk1_0,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk3_0)), inference(rw,[status(thm)],[c_0_176, c_0_177])).
cnf(c_0_186, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk1_0,esk1_0,esk3_0,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk1_0,esk1_0,esk1_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_178, c_0_172]), c_0_179])).
fof(c_0_187, plain, ![X103, X104, X105]:(X103=X104|(X103=X105|k4_xboole_0(k1_enumset1(X103,X104,X105),k1_tarski(X103))=k2_tarski(X104,X105))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t136_enumset1])])])).
cnf(c_0_188, plain, (k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_180])).
cnf(c_0_189, plain, (k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_181])).
cnf(c_0_190, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk4_0,esk4_0,esk1_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0)), inference(rw,[status(thm)],[c_0_138, c_0_149])).
cnf(c_0_191, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0,esk1_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_182, c_0_128]), c_0_138])).
cnf(c_0_192, negated_conjecture, (~r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0)))), inference(rw,[status(thm)],[c_0_183, c_0_149])).
cnf(c_0_193, negated_conjecture, (k7_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk3_0,esk3_0,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk3_0)), inference(rw,[status(thm)],[c_0_184, c_0_177])).
cnf(c_0_194, negated_conjecture, (k7_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk3_0,esk1_0,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk1_0,esk1_0,esk1_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_185, c_0_179]), c_0_186])).
cnf(c_0_195, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k7_enumset1(X1,X2,X3,X3,X3,X4,X5,X6,X7)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_94]), c_0_106]), c_0_83])).
cnf(c_0_196, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X6,X7)=k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_94]), c_0_75]), c_0_83]), c_0_110])).
cnf(c_0_197, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k6_enumset1(X1,X2,X4,X3,X5,X6,X7,X8)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_89]), c_0_106]), c_0_83]), c_0_83])).
cnf(c_0_198, plain, (X1=X2|X1=X3|k4_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X1))=k2_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_187])).
cnf(c_0_199, plain, (k5_xboole_0(X1,k5_xboole_0(X2,X3))=k5_xboole_0(X3,k5_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_188, c_0_189])).
cnf(c_0_200, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk4_0,esk4_0,esk1_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk3_0)), inference(rw,[status(thm)],[c_0_190, c_0_163])).
cnf(c_0_201, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0))), inference(rw,[status(thm)],[c_0_191, c_0_149])).
cnf(c_0_202, negated_conjecture, (~r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk3_0)))), inference(rw,[status(thm)],[c_0_192, c_0_163])).
cnf(c_0_203, negated_conjecture, (k7_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk3_0,esk3_0,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk1_0,esk1_0,esk1_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_193, c_0_179]), c_0_186])).
cnf(c_0_204, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k6_enumset1(X1,X1,X2,X4,X3,X5,X6,X7)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_89]), c_0_97]), c_0_83]), c_0_83])).
cnf(c_0_205, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk1_0,esk1_0,esk1_0,esk4_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_194, c_0_195]), c_0_196]), c_0_123]), c_0_196]), c_0_197])).
cnf(c_0_206, plain, (X1=X3|X1=X2|k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k3_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)))=k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_198, c_0_64]), c_0_44]), c_0_44]), c_0_67]), c_0_45]), c_0_45]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_49])).
cnf(c_0_207, plain, (k5_xboole_0(X1,k5_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_81]), c_0_82])).
cnf(c_0_208, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0,esk1_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk3_0)), inference(rw,[status(thm)],[c_0_149, c_0_163])).
cnf(c_0_209, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk4_0,esk4_0,esk1_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk1_0,esk1_0,esk1_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_200, c_0_179]), c_0_186])).
cnf(c_0_210, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk3_0))), inference(rw,[status(thm)],[c_0_201, c_0_163])).
cnf(c_0_211, negated_conjecture, (~r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk1_0,esk1_0,esk1_0,esk4_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_202, c_0_179]), c_0_186])).
cnf(c_0_212, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_203, c_0_195]), c_0_196]), c_0_123]), c_0_196]), c_0_196]), c_0_204]), c_0_196]), c_0_196]), c_0_205])).
cnf(c_0_213, plain, (k6_enumset1(X1,X1,X1,X1,X1,X2,X2,X2)=k5_xboole_0(k6_enumset1(X3,X3,X3,X3,X3,X3,X1,X2),k3_xboole_0(k6_enumset1(X3,X3,X3,X3,X3,X3,X1,X2),k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X3)))|X3=X1|X3=X2), inference(spm,[status(thm)],[c_0_94, c_0_206])).
cnf(c_0_214, plain, (k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5),k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9))=k5_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X7,X8,X9),k3_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X7,X8,X9),k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5)))), inference(spm,[status(thm)],[c_0_207, c_0_106])).
cnf(c_0_215, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0,esk1_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk1_0,esk1_0,esk1_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_208, c_0_179]), c_0_186])).
cnf(c_0_216, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk4_0,esk1_0,X1)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk1_0,esk1_0,esk4_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_209]), c_0_75]), c_0_110]), c_0_110])).
cnf(c_0_217, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk1_0,esk1_0,esk1_0,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_210, c_0_179]), c_0_186])).
cnf(c_0_218, negated_conjecture, (~r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk4_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_211, c_0_205]), c_0_212])).
cnf(c_0_219, plain, (k6_enumset1(X1,X1,X1,X1,X1,X2,X2,X2)=k5_xboole_0(k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X3),k6_enumset1(X3,X3,X3,X3,X3,X3,X1,X2))|X3=X2|X3=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_213, c_0_214]), c_0_83])).
cnf(c_0_220, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,esk4_0,X1)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk1_0,X1)), inference(rw,[status(thm)],[c_0_102, c_0_118])).
cnf(c_0_221, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk1_0,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk4_0,esk4_0,esk4_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_115, c_0_118]), c_0_118])).
cnf(c_0_222, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk1_0,esk1_0,esk4_0,esk1_0)=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk1_0,esk1_0,esk1_0,esk4_0)), inference(rw,[status(thm)],[c_0_215, c_0_216])).
cnf(c_0_223, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_217, c_0_205]), c_0_212])).
cnf(c_0_224, negated_conjecture, (esk1_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_225, negated_conjecture, (esk1_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_226, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_219]), c_0_220]), c_0_221]), c_0_124]), c_0_126]), c_0_132]), c_0_216]), c_0_222]), c_0_205]), c_0_212]), c_0_223])]), c_0_224]), c_0_225]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 227
# Proof object clause steps            : 162
# Proof object formula steps           : 65
# Proof object conjectures             : 98
# Proof object clause conjectures      : 95
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 34
# Proof object initial formulas used   : 32
# Proof object generating inferences   : 51
# Proof object simplifying inferences  : 373
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 32
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 25
# Processed clauses                    : 4956
# ...of these trivial                  : 373
# ...subsumed                          : 3767
# ...remaining for further processing  : 816
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 400
# Generated clauses                    : 46509
# ...of the previous two non-trivial   : 35694
# Contextual simplify-reflections      : 0
# Paramodulations                      : 46509
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 391
#    Positive orientable unit clauses  : 205
#    Positive unorientable unit clauses: 36
#    Negative unit clauses             : 111
#    Non-unit-clauses                  : 39
# Current number of unprocessed clauses: 29500
# ...number of literals in the above   : 57281
# Current number of archived formulas  : 0
# Current number of archived clauses   : 434
# Clause-clause subsumption calls (NU) : 1016
# Rec. Clause-clause subsumption calls : 533
# Non-unit clause-clause subsumptions  : 204
# Unit Clause-clause subsumption calls : 8455
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8314
# BW rewrite match successes           : 2509
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 730450

# -------------------------------------------------
# User time                : 0.322 s
# System time              : 0.010 s
# Total time               : 0.332 s
# Maximum resident set size: 27160 pages
