# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t72_zfmisc_1, axiom, ![X1, X2, X3]:(k4_xboole_0(k2_tarski(X1,X2),X3)=k2_tarski(X1,X2)<=>(~(r2_hidden(X1,X3))&~(r2_hidden(X2,X3)))), file('data/MPTP/proofs/490e706dbfc74fac8604d1b9466f4b95.E_input', t72_zfmisc_1)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/490e706dbfc74fac8604d1b9466f4b95.E_input', redefinition_k6_subset_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/490e706dbfc74fac8604d1b9466f4b95.E_input', t3_boole)).
fof(t29_yellow_6, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>r1_waybel_0(X1,X2,u1_struct_0(X1)))), file('data/MPTP/proofs/490e706dbfc74fac8604d1b9466f4b95.E_input', t29_yellow_6)).
fof(t10_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r2_waybel_0(X1,X2,X3)<=>~(r1_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3)))))), file('data/MPTP/proofs/490e706dbfc74fac8604d1b9466f4b95.E_input', t10_waybel_0)).
fof(d12_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r2_waybel_0(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X4,X5)&r2_hidden(k2_waybel_0(X1,X2,X5),X3))))))), file('data/MPTP/proofs/490e706dbfc74fac8604d1b9466f4b95.E_input', d12_waybel_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/490e706dbfc74fac8604d1b9466f4b95.E_input', d3_tarski)).
fof(existence_m1_subset_1, axiom, ![X1]:?[X2]:m1_subset_1(X2,X1), file('data/MPTP/proofs/490e706dbfc74fac8604d1b9466f4b95.E_input', existence_m1_subset_1)).
fof(t8_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3, X4]:(r1_tarski(X3,X4)=>((r1_waybel_0(X1,X2,X3)=>r1_waybel_0(X1,X2,X4))&(r2_waybel_0(X1,X2,X3)=>r2_waybel_0(X1,X2,X4)))))), file('data/MPTP/proofs/490e706dbfc74fac8604d1b9466f4b95.E_input', t8_waybel_0)).
fof(c_0_9, plain, ![X32, X33, X34]:(((~r2_hidden(X32,X34)|k4_xboole_0(k2_tarski(X32,X33),X34)!=k2_tarski(X32,X33))&(~r2_hidden(X33,X34)|k4_xboole_0(k2_tarski(X32,X33),X34)!=k2_tarski(X32,X33)))&(r2_hidden(X32,X34)|r2_hidden(X33,X34)|k4_xboole_0(k2_tarski(X32,X33),X34)=k2_tarski(X32,X33))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t72_zfmisc_1])])])])).
fof(c_0_10, plain, ![X17, X18]:k6_subset_1(X17,X18)=k4_xboole_0(X17,X18), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
fof(c_0_11, plain, ![X19]:k4_xboole_0(X19,k1_xboole_0)=X19, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_12, plain, (~r2_hidden(X1,X2)|k4_xboole_0(k2_tarski(X3,X1),X2)!=k2_tarski(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_15, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>r1_waybel_0(X1,X2,u1_struct_0(X1))))), inference(assume_negation,[status(cth)],[t29_yellow_6])).
fof(c_0_16, plain, ![X14, X15, X16]:((~r2_waybel_0(X14,X15,X16)|~r1_waybel_0(X14,X15,k6_subset_1(u1_struct_0(X14),X16))|(v2_struct_0(X15)|~l1_waybel_0(X15,X14))|(v2_struct_0(X14)|~l1_struct_0(X14)))&(r1_waybel_0(X14,X15,k6_subset_1(u1_struct_0(X14),X16))|r2_waybel_0(X14,X15,X16)|(v2_struct_0(X15)|~l1_waybel_0(X15,X14))|(v2_struct_0(X14)|~l1_struct_0(X14)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t10_waybel_0])])])])])).
cnf(c_0_17, plain, (k6_subset_1(k2_tarski(X3,X1),X2)!=k2_tarski(X3,X1)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_18, plain, (k6_subset_1(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_14, c_0_13])).
fof(c_0_19, plain, ![X20, X21, X22, X23, X25, X27]:(((m1_subset_1(esk4_4(X20,X21,X22,X23),u1_struct_0(X21))|~m1_subset_1(X23,u1_struct_0(X21))|~r2_waybel_0(X20,X21,X22)|(v2_struct_0(X21)|~l1_waybel_0(X21,X20))|(v2_struct_0(X20)|~l1_struct_0(X20)))&((r1_orders_2(X21,X23,esk4_4(X20,X21,X22,X23))|~m1_subset_1(X23,u1_struct_0(X21))|~r2_waybel_0(X20,X21,X22)|(v2_struct_0(X21)|~l1_waybel_0(X21,X20))|(v2_struct_0(X20)|~l1_struct_0(X20)))&(r2_hidden(k2_waybel_0(X20,X21,esk4_4(X20,X21,X22,X23)),X22)|~m1_subset_1(X23,u1_struct_0(X21))|~r2_waybel_0(X20,X21,X22)|(v2_struct_0(X21)|~l1_waybel_0(X21,X20))|(v2_struct_0(X20)|~l1_struct_0(X20)))))&((m1_subset_1(esk5_3(X20,X21,X25),u1_struct_0(X21))|r2_waybel_0(X20,X21,X25)|(v2_struct_0(X21)|~l1_waybel_0(X21,X20))|(v2_struct_0(X20)|~l1_struct_0(X20)))&(~m1_subset_1(X27,u1_struct_0(X21))|(~r1_orders_2(X21,esk5_3(X20,X21,X25),X27)|~r2_hidden(k2_waybel_0(X20,X21,X27),X25))|r2_waybel_0(X20,X21,X25)|(v2_struct_0(X21)|~l1_waybel_0(X21,X20))|(v2_struct_0(X20)|~l1_struct_0(X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d12_waybel_0])])])])])])])).
fof(c_0_20, negated_conjecture, ((~v2_struct_0(esk1_0)&l1_struct_0(esk1_0))&((~v2_struct_0(esk2_0)&(v4_orders_2(esk2_0)&(v7_waybel_0(esk2_0)&l1_waybel_0(esk2_0,esk1_0))))&~r1_waybel_0(esk1_0,esk2_0,u1_struct_0(esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_15])])])])).
cnf(c_0_21, plain, (r1_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3))|r2_waybel_0(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, plain, ![X8, X9, X10, X11, X12]:((~r1_tarski(X8,X9)|(~r2_hidden(X10,X8)|r2_hidden(X10,X9)))&((r2_hidden(esk3_2(X11,X12),X11)|r1_tarski(X11,X12))&(~r2_hidden(esk3_2(X11,X12),X12)|r1_tarski(X11,X12)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_23, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_24, plain, (r2_hidden(k2_waybel_0(X1,X2,esk4_4(X1,X2,X3,X4)),X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~r2_waybel_0(X1,X2,X3)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_25, plain, ![X35]:m1_subset_1(esk6_1(X35),X35), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[existence_m1_subset_1])])).
fof(c_0_26, plain, ![X28, X29, X30, X31]:((~r1_waybel_0(X28,X29,X30)|r1_waybel_0(X28,X29,X31)|~r1_tarski(X30,X31)|(v2_struct_0(X29)|~l1_waybel_0(X29,X28))|(v2_struct_0(X28)|~l1_struct_0(X28)))&(~r2_waybel_0(X28,X29,X30)|r2_waybel_0(X28,X29,X31)|~r1_tarski(X30,X31)|(v2_struct_0(X29)|~l1_waybel_0(X29,X28))|(v2_struct_0(X28)|~l1_struct_0(X28)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_waybel_0])])])])])).
cnf(c_0_27, negated_conjecture, (~r1_waybel_0(esk1_0,esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (r2_waybel_0(X1,X2,k1_xboole_0)|r1_waybel_0(X1,X2,u1_struct_0(X1))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_21, c_0_18])).
cnf(c_0_29, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30, negated_conjecture, (l1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_31, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_32, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_33, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_34, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_waybel_0(X1,X2,k1_xboole_0)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_35, plain, (m1_subset_1(esk6_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_36, plain, (r2_waybel_0(X1,X2,X4)|v2_struct_0(X2)|v2_struct_0(X1)|~r2_waybel_0(X1,X2,X3)|~r1_tarski(X3,X4)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_37, negated_conjecture, (r2_waybel_0(esk1_0,esk2_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30])]), c_0_31]), c_0_32])).
cnf(c_0_38, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_23, c_0_33])).
cnf(c_0_39, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X2,X1,k1_xboole_0)|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40, negated_conjecture, (r2_waybel_0(esk1_0,esk2_0,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_29]), c_0_30])]), c_0_32]), c_0_31])).
cnf(c_0_41, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_29]), c_0_30])]), c_0_32]), c_0_31]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 42
# Proof object clause steps            : 23
# Proof object formula steps           : 19
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 24
# Processed clauses                    : 64
# ...of these trivial                  : 0
# ...subsumed                          : 5
# ...remaining for further processing  : 59
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 1
# Generated clauses                    : 26
# ...of the previous two non-trivial   : 19
# Contextual simplify-reflections      : 0
# Paramodulations                      : 26
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 32
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 19
# Current number of unprocessed clauses: 3
# ...number of literals in the above   : 31
# Current number of archived formulas  : 0
# Current number of archived clauses   : 28
# Clause-clause subsumption calls (NU) : 160
# Rec. Clause-clause subsumption calls : 25
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 16
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2629

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.003 s
# Total time               : 0.010 s
# Maximum resident set size: 28760 pages
