# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/494cf13896c24deab49a8ee4b83010e9.E_input', d3_tarski)).
fof(l49_zfmisc_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>r1_tarski(X1,k3_tarski(X2))), file('data/MPTP/proofs/494cf13896c24deab49a8ee4b83010e9.E_input', l49_zfmisc_1)).
fof(d2_setfam_1, axiom, ![X1, X2]:(r1_setfam_1(X1,X2)<=>![X3]:~((r2_hidden(X3,X1)&![X4]:~((r2_hidden(X4,X2)&r1_tarski(X3,X4)))))), file('data/MPTP/proofs/494cf13896c24deab49a8ee4b83010e9.E_input', d2_setfam_1)).
fof(t18_setfam_1, conjecture, ![X1, X2]:(r1_setfam_1(X1,X2)=>r1_tarski(k3_tarski(X1),k3_tarski(X2))), file('data/MPTP/proofs/494cf13896c24deab49a8ee4b83010e9.E_input', t18_setfam_1)).
fof(t94_zfmisc_1, axiom, ![X1, X2]:(![X3]:(r2_hidden(X3,X1)=>r1_tarski(X3,X2))=>r1_tarski(k3_tarski(X1),X2)), file('data/MPTP/proofs/494cf13896c24deab49a8ee4b83010e9.E_input', t94_zfmisc_1)).
fof(c_0_5, plain, ![X5, X6, X7, X8, X9]:((~r1_tarski(X5,X6)|(~r2_hidden(X7,X5)|r2_hidden(X7,X6)))&((r2_hidden(esk1_2(X8,X9),X8)|r1_tarski(X8,X9))&(~r2_hidden(esk1_2(X8,X9),X9)|r1_tarski(X8,X9)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_6, plain, ![X11, X12]:(~r2_hidden(X11,X12)|r1_tarski(X11,k3_tarski(X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l49_zfmisc_1])])).
fof(c_0_7, plain, ![X13, X14, X15, X17, X18, X20]:(((r2_hidden(esk2_3(X13,X14,X15),X14)|~r2_hidden(X15,X13)|~r1_setfam_1(X13,X14))&(r1_tarski(X15,esk2_3(X13,X14,X15))|~r2_hidden(X15,X13)|~r1_setfam_1(X13,X14)))&((r2_hidden(esk3_2(X17,X18),X17)|r1_setfam_1(X17,X18))&(~r2_hidden(X20,X18)|~r1_tarski(esk3_2(X17,X18),X20)|r1_setfam_1(X17,X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_setfam_1])])])])])])).
cnf(c_0_8, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (r1_tarski(X1,esk2_3(X2,X3,X1))|~r2_hidden(X1,X2)|~r1_setfam_1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (r2_hidden(X1,k3_tarski(X2))|~r2_hidden(X1,X3)|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_12, plain, (r2_hidden(X1,esk2_3(X2,X3,X4))|~r1_setfam_1(X2,X3)|~r2_hidden(X1,X4)|~r2_hidden(X4,X2)), inference(spm,[status(thm)],[c_0_8, c_0_10])).
fof(c_0_13, negated_conjecture, ~(![X1, X2]:(r1_setfam_1(X1,X2)=>r1_tarski(k3_tarski(X1),k3_tarski(X2)))), inference(assume_negation,[status(cth)],[t18_setfam_1])).
cnf(c_0_14, plain, (r2_hidden(X1,k3_tarski(X2))|~r1_setfam_1(X3,X4)|~r2_hidden(esk2_3(X3,X4,X5),X2)|~r2_hidden(X1,X5)|~r2_hidden(X5,X3)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_15, plain, (r2_hidden(esk2_3(X1,X2,X3),X2)|~r2_hidden(X3,X1)|~r1_setfam_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_16, negated_conjecture, (r1_setfam_1(esk4_0,esk5_0)&~r1_tarski(k3_tarski(esk4_0),k3_tarski(esk5_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
cnf(c_0_17, plain, (r2_hidden(X1,k3_tarski(X2))|~r1_setfam_1(X3,X2)|~r2_hidden(X1,X4)|~r2_hidden(X4,X3)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_18, negated_conjecture, (r1_setfam_1(esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_19, plain, ![X23, X24]:((r2_hidden(esk6_2(X23,X24),X23)|r1_tarski(k3_tarski(X23),X24))&(~r1_tarski(esk6_2(X23,X24),X24)|r1_tarski(k3_tarski(X23),X24))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t94_zfmisc_1])])])])).
cnf(c_0_20, negated_conjecture, (r2_hidden(X1,k3_tarski(esk5_0))|~r2_hidden(X2,esk4_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_21, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_tarski(k3_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_22, negated_conjecture, (r2_hidden(X1,k3_tarski(esk5_0))|r1_tarski(k3_tarski(esk4_0),X2)|~r2_hidden(X1,esk6_2(esk4_0,X2))), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_23, plain, (r2_hidden(esk1_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_24, plain, (r1_tarski(X1,X2)|~r2_hidden(esk1_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_25, negated_conjecture, (r2_hidden(esk1_2(esk6_2(esk4_0,X1),X2),k3_tarski(esk5_0))|r1_tarski(esk6_2(esk4_0,X1),X2)|r1_tarski(k3_tarski(esk4_0),X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_26, plain, (r1_tarski(k3_tarski(X1),X2)|~r1_tarski(esk6_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, negated_conjecture, (r1_tarski(esk6_2(esk4_0,X1),k3_tarski(esk5_0))|r1_tarski(k3_tarski(esk4_0),X1)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_28, negated_conjecture, (~r1_tarski(k3_tarski(esk4_0),k3_tarski(esk5_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_29, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 30
# Proof object clause steps            : 19
# Proof object formula steps           : 11
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 1
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 12
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 12
# Processed clauses                    : 282
# ...of these trivial                  : 0
# ...subsumed                          : 52
# ...remaining for further processing  : 230
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 1089
# ...of the previous two non-trivial   : 1085
# Contextual simplify-reflections      : 0
# Paramodulations                      : 1089
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 218
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 213
# Current number of unprocessed clauses: 826
# ...number of literals in the above   : 4265
# Current number of archived formulas  : 0
# Current number of archived clauses   : 12
# Clause-clause subsumption calls (NU) : 5152
# Rec. Clause-clause subsumption calls : 2842
# Non-unit clause-clause subsumptions  : 52
# Unit Clause-clause subsumption calls : 8
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 22312

# -------------------------------------------------
# User time                : 0.026 s
# System time              : 0.000 s
# Total time               : 0.026 s
# Maximum resident set size: 28616 pages
