# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_connsp_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(m1_connsp_2(X3,X1,X2)<=>r2_hidden(X2,k1_tops_1(X1,X3)))))), file('data/MPTP/proofs/4955035a70d34bf8bca12e31d8f797e1.E_input', d1_connsp_2)).
fof(dt_m1_connsp_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,u1_struct_0(X1)))=>![X3]:(m1_connsp_2(X3,X1,X2)=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/4955035a70d34bf8bca12e31d8f797e1.E_input', dt_m1_connsp_2)).
fof(t4_yellow19, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((v13_waybel_0(X3,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))))))=>(r2_waybel_7(X1,X3,X2)<=>r1_tarski(k1_yellow19(X1,X2),X3))))), file('data/MPTP/proofs/4955035a70d34bf8bca12e31d8f797e1.E_input', t4_yellow19)).
fof(t5_connsp_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((v3_pre_topc(X2,X1)&r2_hidden(X3,X2))=>m1_connsp_2(X2,X1,X3))))), file('data/MPTP/proofs/4955035a70d34bf8bca12e31d8f797e1.E_input', t5_connsp_2)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/4955035a70d34bf8bca12e31d8f797e1.E_input', t4_subset)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/4955035a70d34bf8bca12e31d8f797e1.E_input', d3_tarski)).
fof(t44_tops_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>r1_tarski(k1_tops_1(X1,X2),X2))), file('data/MPTP/proofs/4955035a70d34bf8bca12e31d8f797e1.E_input', t44_tops_1)).
fof(d5_waybel_7, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2, X3]:(r2_waybel_7(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>((v3_pre_topc(X4,X1)&r2_hidden(X3,X4))=>r2_hidden(X4,X2))))), file('data/MPTP/proofs/4955035a70d34bf8bca12e31d8f797e1.E_input', d5_waybel_7)).
fof(dt_k1_tops_1, axiom, ![X1, X2]:((l1_pre_topc(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k1_tops_1(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/4955035a70d34bf8bca12e31d8f797e1.E_input', dt_k1_tops_1)).
fof(fc9_tops_1, axiom, ![X1, X2]:(((v2_pre_topc(X1)&l1_pre_topc(X1))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>v3_pre_topc(k1_tops_1(X1,X2),X1)), file('data/MPTP/proofs/4955035a70d34bf8bca12e31d8f797e1.E_input', fc9_tops_1)).
fof(t3_yellow19, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r2_hidden(X3,k1_yellow19(X1,X2))<=>m1_connsp_2(X3,X1,X2)))), file('data/MPTP/proofs/4955035a70d34bf8bca12e31d8f797e1.E_input', t3_yellow19)).
fof(t11_waybel_7, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X1))))=>(v13_waybel_0(X2,k3_yellow_1(X1))<=>![X3, X4]:((r1_tarski(X3,X4)&(r1_tarski(X4,X1)&r2_hidden(X3,X2)))=>r2_hidden(X4,X2)))), file('data/MPTP/proofs/4955035a70d34bf8bca12e31d8f797e1.E_input', t11_waybel_7)).
fof(d2_yellow_1, axiom, ![X1]:k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1)), file('data/MPTP/proofs/4955035a70d34bf8bca12e31d8f797e1.E_input', d2_yellow_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/4955035a70d34bf8bca12e31d8f797e1.E_input', t3_subset)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/4955035a70d34bf8bca12e31d8f797e1.E_input', d3_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/4955035a70d34bf8bca12e31d8f797e1.E_input', dt_l1_pre_topc)).
fof(c_0_16, plain, ![X5, X6, X7]:((~m1_connsp_2(X7,X5,X6)|r2_hidden(X6,k1_tops_1(X5,X7))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5)))|~m1_subset_1(X6,u1_struct_0(X5))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))&(~r2_hidden(X6,k1_tops_1(X5,X7))|m1_connsp_2(X7,X5,X6)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5)))|~m1_subset_1(X6,u1_struct_0(X5))|(v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_connsp_2])])])])])).
fof(c_0_17, plain, ![X8, X9, X10]:(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))|~m1_subset_1(X9,u1_struct_0(X8))|(~m1_connsp_2(X10,X8,X9)|m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X8))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_connsp_2])])])])).
fof(c_0_18, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((v13_waybel_0(X3,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))))))=>(r2_waybel_7(X1,X3,X2)<=>r1_tarski(k1_yellow19(X1,X2),X3)))))), inference(assume_negation,[status(cth)],[t4_yellow19])).
fof(c_0_19, plain, ![X14, X15, X16]:(v2_struct_0(X14)|(~v2_pre_topc(X14)|~l1_pre_topc(X14))|(~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X14)))|(~m1_subset_1(X16,u1_struct_0(X14))|(~v3_pre_topc(X15,X14)|~r2_hidden(X16,X15)|m1_connsp_2(X15,X14,X16))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t5_connsp_2])])])])).
fof(c_0_20, plain, ![X17, X18, X19]:(~r2_hidden(X17,X18)|~m1_subset_1(X18,k1_zfmisc_1(X19))|m1_subset_1(X17,X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_21, plain, ![X20, X21, X22, X23, X24]:((~r1_tarski(X20,X21)|(~r2_hidden(X22,X20)|r2_hidden(X22,X21)))&((r2_hidden(esk4_2(X23,X24),X23)|r1_tarski(X23,X24))&(~r2_hidden(esk4_2(X23,X24),X24)|r1_tarski(X23,X24)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_22, plain, ![X26, X27]:(~l1_pre_topc(X26)|(~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X26)))|r1_tarski(k1_tops_1(X26,X27),X27))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t44_tops_1])])])).
cnf(c_0_23, plain, (r2_hidden(X3,k1_tops_1(X2,X1))|v2_struct_0(X2)|~m1_connsp_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,u1_struct_0(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (v2_struct_0(X1)|m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_connsp_2(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_25, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&((v13_waybel_0(esk3_0,k3_yellow_1(k2_struct_0(esk1_0)))&m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))))&((~r2_waybel_7(esk1_0,esk3_0,esk2_0)|~r1_tarski(k1_yellow19(esk1_0,esk2_0),esk3_0))&(r2_waybel_7(esk1_0,esk3_0,esk2_0)|r1_tarski(k1_yellow19(esk1_0,esk2_0),esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_18])])])])).
cnf(c_0_26, plain, (v2_struct_0(X1)|m1_connsp_2(X2,X1,X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_pre_topc(X2,X1)|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_28, plain, ![X28, X29, X30, X31, X32, X33]:((~r2_waybel_7(X28,X29,X30)|(~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X28)))|(~v3_pre_topc(X31,X28)|~r2_hidden(X30,X31)|r2_hidden(X31,X29)))|(~v2_pre_topc(X28)|~l1_pre_topc(X28)))&((m1_subset_1(esk5_3(X28,X32,X33),k1_zfmisc_1(u1_struct_0(X28)))|r2_waybel_7(X28,X32,X33)|(~v2_pre_topc(X28)|~l1_pre_topc(X28)))&(((v3_pre_topc(esk5_3(X28,X32,X33),X28)|r2_waybel_7(X28,X32,X33)|(~v2_pre_topc(X28)|~l1_pre_topc(X28)))&(r2_hidden(X33,esk5_3(X28,X32,X33))|r2_waybel_7(X28,X32,X33)|(~v2_pre_topc(X28)|~l1_pre_topc(X28))))&(~r2_hidden(esk5_3(X28,X32,X33),X32)|r2_waybel_7(X28,X32,X33)|(~v2_pre_topc(X28)|~l1_pre_topc(X28)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_waybel_7])])])])])])).
cnf(c_0_29, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, plain, (r1_tarski(k1_tops_1(X1,X2),X2)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, plain, (r2_hidden(X1,k1_tops_1(X2,X3))|v2_struct_0(X2)|~m1_connsp_2(X3,X2,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_32, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_36, plain, (m1_connsp_2(X1,X2,X3)|v2_struct_0(X2)|~v3_pre_topc(X1,X2)|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_37, plain, (m1_subset_1(esk5_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|r2_waybel_7(X1,X2,X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, plain, (v3_pre_topc(esk5_3(X1,X2,X3),X1)|r2_waybel_7(X1,X2,X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_39, plain, ![X38, X39]:(~l1_pre_topc(X38)|~m1_subset_1(X39,k1_zfmisc_1(u1_struct_0(X38)))|m1_subset_1(k1_tops_1(X38,X39),k1_zfmisc_1(u1_struct_0(X38)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_tops_1])])).
fof(c_0_40, plain, ![X40, X41]:(~v2_pre_topc(X40)|~l1_pre_topc(X40)|~m1_subset_1(X41,k1_zfmisc_1(u1_struct_0(X40)))|v3_pre_topc(k1_tops_1(X40,X41),X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc9_tops_1])])).
fof(c_0_41, plain, ![X35, X36, X37]:((~r2_hidden(X37,k1_yellow19(X35,X36))|m1_connsp_2(X37,X35,X36)|~m1_subset_1(X36,u1_struct_0(X35))|(v2_struct_0(X35)|(~v2_pre_topc(X35)|~l1_pre_topc(X35))))&(~m1_connsp_2(X37,X35,X36)|r2_hidden(X37,k1_yellow19(X35,X36))|~m1_subset_1(X36,u1_struct_0(X35))|(v2_struct_0(X35)|(~v2_pre_topc(X35)|~l1_pre_topc(X35))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_yellow19])])])])])).
cnf(c_0_42, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_tops_1(X3,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_43, negated_conjecture, (r2_hidden(esk2_0,k1_tops_1(esk1_0,X1))|~m1_connsp_2(X1,esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_44, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_connsp_2(X1,esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_32]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_45, plain, (r2_waybel_7(X1,X2,X3)|m1_connsp_2(esk5_3(X1,X2,X3),X1,X4)|v2_struct_0(X1)|~r2_hidden(X4,esk5_3(X1,X2,X3))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_46, plain, (r2_hidden(X1,esk5_3(X2,X3,X1))|r2_waybel_7(X2,X3,X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_47, plain, (m1_subset_1(k1_tops_1(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_48, plain, (v3_pre_topc(k1_tops_1(X1,X2),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_49, negated_conjecture, (r2_waybel_7(esk1_0,esk3_0,esk2_0)|r1_tarski(k1_yellow19(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_50, plain, (r2_hidden(X1,k1_yellow19(X2,X3))|v2_struct_0(X2)|~m1_connsp_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_51, plain, (r2_hidden(X4,X2)|~r2_waybel_7(X1,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v3_pre_topc(X4,X1)|~r2_hidden(X3,X4)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk2_0,X1)|~m1_connsp_2(X1,esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_33])]), c_0_44])).
cnf(c_0_53, plain, (r2_waybel_7(X1,X2,X3)|m1_connsp_2(esk5_3(X1,X2,X3),X1,X3)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_54, plain, (m1_connsp_2(k1_tops_1(X1,X2),X1,X3)|v2_struct_0(X1)|~r2_hidden(X3,k1_tops_1(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_47]), c_0_48])).
cnf(c_0_55, negated_conjecture, (r2_waybel_7(esk1_0,esk3_0,esk2_0)|r2_hidden(X1,esk3_0)|~r2_hidden(X1,k1_yellow19(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_29, c_0_49])).
cnf(c_0_56, negated_conjecture, (r2_hidden(X1,k1_yellow19(esk1_0,esk2_0))|~m1_connsp_2(X1,esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_32]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_57, plain, (r2_waybel_7(X1,X2,X3)|r2_hidden(esk5_3(X1,X2,X3),X4)|~r2_waybel_7(X1,X4,X5)|~r2_hidden(X5,esk5_3(X1,X2,X3))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_37]), c_0_38])).
cnf(c_0_58, negated_conjecture, (r2_waybel_7(esk1_0,X1,esk2_0)|r2_hidden(esk2_0,esk5_3(esk1_0,X1,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_59, negated_conjecture, (m1_connsp_2(k1_tops_1(esk1_0,X1),esk1_0,esk2_0)|~m1_connsp_2(X1,esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_43]), c_0_33]), c_0_34])]), c_0_35]), c_0_44])).
cnf(c_0_60, negated_conjecture, (r2_waybel_7(esk1_0,esk3_0,esk2_0)|r2_hidden(X1,esk3_0)|~m1_connsp_2(X1,esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_61, negated_conjecture, (r2_waybel_7(esk1_0,X1,esk2_0)|r2_hidden(esk5_3(esk1_0,X1,esk2_0),X2)|~r2_waybel_7(esk1_0,X2,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_33]), c_0_34])])).
fof(c_0_62, plain, ![X47, X48, X49, X50]:((~v13_waybel_0(X48,k3_yellow_1(X47))|(~r1_tarski(X49,X50)|(~r1_tarski(X50,X47)|~r2_hidden(X49,X48))|r2_hidden(X50,X48))|~m1_subset_1(X48,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X47)))))&(((r1_tarski(esk6_2(X47,X48),esk7_2(X47,X48))|v13_waybel_0(X48,k3_yellow_1(X47))|~m1_subset_1(X48,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X47)))))&((r1_tarski(esk7_2(X47,X48),X47)|v13_waybel_0(X48,k3_yellow_1(X47))|~m1_subset_1(X48,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X47)))))&(r2_hidden(esk6_2(X47,X48),X48)|v13_waybel_0(X48,k3_yellow_1(X47))|~m1_subset_1(X48,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X47)))))))&(~r2_hidden(esk7_2(X47,X48),X48)|v13_waybel_0(X48,k3_yellow_1(X47))|~m1_subset_1(X48,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X47))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_waybel_7])])])])])).
fof(c_0_63, plain, ![X42]:k3_yellow_1(X42)=k3_lattice3(k1_lattice3(X42)), inference(variable_rename,[status(thm)],[d2_yellow_1])).
cnf(c_0_64, negated_conjecture, (m1_subset_1(k1_tops_1(esk1_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_connsp_2(X1,esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_44, c_0_59])).
cnf(c_0_65, plain, (r2_waybel_7(X1,X2,X3)|~r2_hidden(esk5_3(X1,X2,X3),X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_66, negated_conjecture, (r2_waybel_7(esk1_0,X1,esk2_0)|r2_hidden(esk5_3(esk1_0,X1,esk2_0),esk3_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_53]), c_0_33]), c_0_34])]), c_0_35]), c_0_61])).
cnf(c_0_67, plain, (r2_hidden(X4,X1)|~v13_waybel_0(X1,k3_yellow_1(X2))|~r1_tarski(X3,X4)|~r1_tarski(X4,X2)|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_68, plain, (k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_69, negated_conjecture, (v13_waybel_0(esk3_0,k3_yellow_1(k2_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_70, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_71, negated_conjecture, (r2_hidden(k1_tops_1(esk1_0,X1),X2)|~v3_pre_topc(k1_tops_1(esk1_0,X1),esk1_0)|~r2_waybel_7(esk1_0,X2,X3)|~r2_hidden(X3,k1_tops_1(esk1_0,X1))|~m1_connsp_2(X1,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_64]), c_0_33]), c_0_34])])).
cnf(c_0_72, negated_conjecture, (r2_waybel_7(esk1_0,esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_33]), c_0_34])])).
cnf(c_0_73, plain, (r2_hidden(X4,X1)|~r2_hidden(X3,X1)|~r1_tarski(X4,X2)|~r1_tarski(X3,X4)|~v13_waybel_0(X1,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_68]), c_0_68])).
cnf(c_0_74, negated_conjecture, (v13_waybel_0(esk3_0,k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))), inference(rw,[status(thm)],[c_0_69, c_0_68])).
cnf(c_0_75, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))))), inference(rw,[status(thm)],[c_0_70, c_0_68])).
cnf(c_0_76, negated_conjecture, (r2_hidden(k1_tops_1(esk1_0,X1),esk3_0)|~v3_pre_topc(k1_tops_1(esk1_0,X1),esk1_0)|~m1_connsp_2(X1,esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_43])).
cnf(c_0_77, negated_conjecture, (r2_hidden(X1,esk3_0)|~r1_tarski(X1,k2_struct_0(esk1_0))|~r1_tarski(X2,X1)|~r2_hidden(X2,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_75])])).
cnf(c_0_78, negated_conjecture, (r2_hidden(k1_tops_1(esk1_0,X1),esk3_0)|~m1_connsp_2(X1,esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_48]), c_0_33]), c_0_34])]), c_0_44])).
fof(c_0_79, plain, ![X45, X46]:((~m1_subset_1(X45,k1_zfmisc_1(X46))|r1_tarski(X45,X46))&(~r1_tarski(X45,X46)|m1_subset_1(X45,k1_zfmisc_1(X46)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_80, negated_conjecture, (r2_hidden(X1,esk3_0)|~r1_tarski(k1_tops_1(esk1_0,X2),X1)|~r1_tarski(X1,k2_struct_0(esk1_0))|~m1_connsp_2(X2,esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_81, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_79])).
fof(c_0_82, plain, ![X43]:(~l1_struct_0(X43)|k2_struct_0(X43)=u1_struct_0(X43)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
fof(c_0_83, plain, ![X44]:(~l1_pre_topc(X44)|l1_struct_0(X44)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_84, negated_conjecture, (r2_hidden(X1,esk3_0)|~r1_tarski(k1_tops_1(esk1_0,X2),X1)|~m1_connsp_2(X2,esk1_0,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_85, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_86, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_87, plain, (m1_connsp_2(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,k1_yellow19(X2,X3))|~m1_subset_1(X3,u1_struct_0(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_88, negated_conjecture, (r2_hidden(X1,esk3_0)|~m1_connsp_2(X1,esk1_0,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_30]), c_0_33])]), c_0_44])).
cnf(c_0_89, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_90, negated_conjecture, (m1_connsp_2(X1,esk1_0,esk2_0)|~r2_hidden(X1,k1_yellow19(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_32]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_91, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_92, negated_conjecture, (r2_hidden(X1,esk3_0)|~m1_connsp_2(X1,esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_33])]), c_0_44])).
cnf(c_0_93, negated_conjecture, (r1_tarski(k1_yellow19(esk1_0,esk2_0),X1)|m1_connsp_2(esk4_2(k1_yellow19(esk1_0,esk2_0),X1),esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_94, negated_conjecture, (~r2_waybel_7(esk1_0,esk3_0,esk2_0)|~r1_tarski(k1_yellow19(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_95, plain, (r1_tarski(X1,X2)|~r2_hidden(esk4_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_96, negated_conjecture, (r1_tarski(k1_yellow19(esk1_0,esk2_0),X1)|r2_hidden(esk4_2(k1_yellow19(esk1_0,esk2_0),X1),esk3_0)), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_97, negated_conjecture, (~r1_tarski(k1_yellow19(esk1_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_94, c_0_72])])).
cnf(c_0_98, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_97]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 99
# Proof object clause steps            : 66
# Proof object formula steps           : 33
# Proof object conjectures             : 38
# Proof object clause conjectures      : 35
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 30
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 30
# Proof object simplifying inferences  : 67
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 36
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 35
# Processed clauses                    : 1277
# ...of these trivial                  : 2
# ...subsumed                          : 481
# ...remaining for further processing  : 794
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 74
# Backward-rewritten                   : 412
# Generated clauses                    : 3229
# ...of the previous two non-trivial   : 3153
# Contextual simplify-reflections      : 54
# Paramodulations                      : 3229
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 273
#    Positive orientable unit clauses  : 18
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 252
# Current number of unprocessed clauses: 1808
# ...number of literals in the above   : 9367
# Current number of archived formulas  : 0
# Current number of archived clauses   : 522
# Clause-clause subsumption calls (NU) : 101611
# Rec. Clause-clause subsumption calls : 18275
# Non-unit clause-clause subsumptions  : 608
# Unit Clause-clause subsumption calls : 586
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 17
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 110667

# -------------------------------------------------
# User time                : 0.121 s
# System time              : 0.000 s
# Total time               : 0.121 s
# Maximum resident set size: 28836 pages
