# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t28_connsp_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>~((![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X4,X3)<=>(v3_pre_topc(X4,X1)&(v4_pre_topc(X4,X1)&r2_hidden(X2,X4)))))&X3=k1_xboole_0))))), file('data/MPTP/proofs/499061f05d2c43f7ac4d88d48f8db4e3.E_input', t28_connsp_2)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/499061f05d2c43f7ac4d88d48f8db4e3.E_input', t2_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/499061f05d2c43f7ac4d88d48f8db4e3.E_input', t3_boole)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/499061f05d2c43f7ac4d88d48f8db4e3.E_input', t4_subset_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/499061f05d2c43f7ac4d88d48f8db4e3.E_input', t2_subset)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/499061f05d2c43f7ac4d88d48f8db4e3.E_input', t7_ordinal1)).
fof(d5_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)), file('data/MPTP/proofs/499061f05d2c43f7ac4d88d48f8db4e3.E_input', d5_subset_1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/499061f05d2c43f7ac4d88d48f8db4e3.E_input', fc2_struct_0)).
fof(t30_tops_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>v4_pre_topc(k3_subset_1(u1_struct_0(X1),X2),X1)))), file('data/MPTP/proofs/499061f05d2c43f7ac4d88d48f8db4e3.E_input', t30_tops_1)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/499061f05d2c43f7ac4d88d48f8db4e3.E_input', dt_l1_pre_topc)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/499061f05d2c43f7ac4d88d48f8db4e3.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/499061f05d2c43f7ac4d88d48f8db4e3.E_input', d4_subset_1)).
fof(d1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1))))))))), file('data/MPTP/proofs/499061f05d2c43f7ac4d88d48f8db4e3.E_input', d1_pre_topc)).
fof(t5_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>r2_hidden(k1_xboole_0,u1_pre_topc(X1))), file('data/MPTP/proofs/499061f05d2c43f7ac4d88d48f8db4e3.E_input', t5_pre_topc)).
fof(d5_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>r2_hidden(X2,u1_pre_topc(X1))))), file('data/MPTP/proofs/499061f05d2c43f7ac4d88d48f8db4e3.E_input', d5_pre_topc)).
fof(c_0_15, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>~((![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X4,X3)<=>(v3_pre_topc(X4,X1)&(v4_pre_topc(X4,X1)&r2_hidden(X2,X4)))))&X3=k1_xboole_0)))))), inference(assume_negation,[status(cth)],[t28_connsp_2])).
fof(c_0_16, plain, ![X9]:r1_tarski(k1_xboole_0,X9), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_17, negated_conjecture, ![X8]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))&((((v3_pre_topc(X8,esk1_0)|~r2_hidden(X8,esk3_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(esk1_0))))&((v4_pre_topc(X8,esk1_0)|~r2_hidden(X8,esk3_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(esk1_0))))&(r2_hidden(esk2_0,X8)|~r2_hidden(X8,esk3_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(esk1_0))))))&(~v3_pre_topc(X8,esk1_0)|(~v4_pre_topc(X8,esk1_0)|~r2_hidden(esk2_0,X8))|r2_hidden(X8,esk3_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(esk1_0)))))&esk3_0=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_15])])])])])])).
fof(c_0_18, plain, ![X10]:k4_xboole_0(X10,k1_xboole_0)=X10, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_19, plain, ![X11]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X11)), inference(variable_rename,[status(thm)],[t4_subset_1])).
fof(c_0_20, plain, ![X12, X13]:(~m1_subset_1(X12,X13)|(v1_xboole_0(X13)|r2_hidden(X12,X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_21, plain, ![X14, X15]:(~r2_hidden(X14,X15)|~r1_tarski(X15,X14)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_22, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, negated_conjecture, (esk3_0=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_25, plain, ![X16, X17]:(~m1_subset_1(X17,k1_zfmisc_1(X16))|k3_subset_1(X16,X17)=k4_xboole_0(X16,X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
cnf(c_0_26, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_27, plain, ![X18]:(v2_struct_0(X18)|~l1_struct_0(X18)|~v1_xboole_0(u1_struct_0(X18))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_28, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_30, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, plain, (r1_tarski(esk3_0,X1)), inference(rw,[status(thm)],[c_0_22, c_0_23])).
fof(c_0_32, plain, ![X19, X20]:((~v3_pre_topc(X20,X19)|v4_pre_topc(k3_subset_1(u1_struct_0(X19),X20),X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|~l1_pre_topc(X19))&(~v4_pre_topc(k3_subset_1(u1_struct_0(X19),X20),X19)|v3_pre_topc(X20,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|~l1_pre_topc(X19))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_tops_1])])])])).
cnf(c_0_33, plain, (k4_xboole_0(X1,esk3_0)=X1), inference(rw,[status(thm)],[c_0_24, c_0_23])).
cnf(c_0_34, plain, (k3_subset_1(X2,X1)=k4_xboole_0(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, plain, (m1_subset_1(esk3_0,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_26, c_0_23])).
cnf(c_0_36, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_38, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_39, plain, ![X21]:(~l1_pre_topc(X21)|l1_struct_0(X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_40, plain, ![X22]:m1_subset_1(k2_subset_1(X22),k1_zfmisc_1(X22)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_41, plain, ![X23]:k2_subset_1(X23)=X23, inference(variable_rename,[status(thm)],[d4_subset_1])).
cnf(c_0_42, negated_conjecture, (r2_hidden(X1,esk3_0)|~v3_pre_topc(X1,esk1_0)|~v4_pre_topc(X1,esk1_0)|~r2_hidden(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_43, plain, (~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_44, plain, (v4_pre_topc(k3_subset_1(u1_struct_0(X2),X1),X2)|~v3_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_45, plain, (k3_subset_1(X1,esk3_0)=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_47, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_48, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_49, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_50, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_51, plain, ![X24, X25, X26, X27]:(((r2_hidden(u1_struct_0(X24),u1_pre_topc(X24))|~v2_pre_topc(X24)|~l1_pre_topc(X24))&((~m1_subset_1(X25,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X24))))|(~r1_tarski(X25,u1_pre_topc(X24))|r2_hidden(k5_setfam_1(u1_struct_0(X24),X25),u1_pre_topc(X24)))|~v2_pre_topc(X24)|~l1_pre_topc(X24))&(~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X24)))|(~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X24)))|(~r2_hidden(X26,u1_pre_topc(X24))|~r2_hidden(X27,u1_pre_topc(X24))|r2_hidden(k9_subset_1(u1_struct_0(X24),X26,X27),u1_pre_topc(X24))))|~v2_pre_topc(X24)|~l1_pre_topc(X24))))&(((m1_subset_1(esk5_1(X24),k1_zfmisc_1(u1_struct_0(X24)))|m1_subset_1(esk4_1(X24),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X24))))|~r2_hidden(u1_struct_0(X24),u1_pre_topc(X24))|v2_pre_topc(X24)|~l1_pre_topc(X24))&((m1_subset_1(esk6_1(X24),k1_zfmisc_1(u1_struct_0(X24)))|m1_subset_1(esk4_1(X24),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X24))))|~r2_hidden(u1_struct_0(X24),u1_pre_topc(X24))|v2_pre_topc(X24)|~l1_pre_topc(X24))&(((r2_hidden(esk5_1(X24),u1_pre_topc(X24))|m1_subset_1(esk4_1(X24),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X24))))|~r2_hidden(u1_struct_0(X24),u1_pre_topc(X24))|v2_pre_topc(X24)|~l1_pre_topc(X24))&(r2_hidden(esk6_1(X24),u1_pre_topc(X24))|m1_subset_1(esk4_1(X24),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X24))))|~r2_hidden(u1_struct_0(X24),u1_pre_topc(X24))|v2_pre_topc(X24)|~l1_pre_topc(X24)))&(~r2_hidden(k9_subset_1(u1_struct_0(X24),esk5_1(X24),esk6_1(X24)),u1_pre_topc(X24))|m1_subset_1(esk4_1(X24),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X24))))|~r2_hidden(u1_struct_0(X24),u1_pre_topc(X24))|v2_pre_topc(X24)|~l1_pre_topc(X24)))))&(((m1_subset_1(esk5_1(X24),k1_zfmisc_1(u1_struct_0(X24)))|r1_tarski(esk4_1(X24),u1_pre_topc(X24))|~r2_hidden(u1_struct_0(X24),u1_pre_topc(X24))|v2_pre_topc(X24)|~l1_pre_topc(X24))&((m1_subset_1(esk6_1(X24),k1_zfmisc_1(u1_struct_0(X24)))|r1_tarski(esk4_1(X24),u1_pre_topc(X24))|~r2_hidden(u1_struct_0(X24),u1_pre_topc(X24))|v2_pre_topc(X24)|~l1_pre_topc(X24))&(((r2_hidden(esk5_1(X24),u1_pre_topc(X24))|r1_tarski(esk4_1(X24),u1_pre_topc(X24))|~r2_hidden(u1_struct_0(X24),u1_pre_topc(X24))|v2_pre_topc(X24)|~l1_pre_topc(X24))&(r2_hidden(esk6_1(X24),u1_pre_topc(X24))|r1_tarski(esk4_1(X24),u1_pre_topc(X24))|~r2_hidden(u1_struct_0(X24),u1_pre_topc(X24))|v2_pre_topc(X24)|~l1_pre_topc(X24)))&(~r2_hidden(k9_subset_1(u1_struct_0(X24),esk5_1(X24),esk6_1(X24)),u1_pre_topc(X24))|r1_tarski(esk4_1(X24),u1_pre_topc(X24))|~r2_hidden(u1_struct_0(X24),u1_pre_topc(X24))|v2_pre_topc(X24)|~l1_pre_topc(X24)))))&((m1_subset_1(esk5_1(X24),k1_zfmisc_1(u1_struct_0(X24)))|~r2_hidden(k5_setfam_1(u1_struct_0(X24),esk4_1(X24)),u1_pre_topc(X24))|~r2_hidden(u1_struct_0(X24),u1_pre_topc(X24))|v2_pre_topc(X24)|~l1_pre_topc(X24))&((m1_subset_1(esk6_1(X24),k1_zfmisc_1(u1_struct_0(X24)))|~r2_hidden(k5_setfam_1(u1_struct_0(X24),esk4_1(X24)),u1_pre_topc(X24))|~r2_hidden(u1_struct_0(X24),u1_pre_topc(X24))|v2_pre_topc(X24)|~l1_pre_topc(X24))&(((r2_hidden(esk5_1(X24),u1_pre_topc(X24))|~r2_hidden(k5_setfam_1(u1_struct_0(X24),esk4_1(X24)),u1_pre_topc(X24))|~r2_hidden(u1_struct_0(X24),u1_pre_topc(X24))|v2_pre_topc(X24)|~l1_pre_topc(X24))&(r2_hidden(esk6_1(X24),u1_pre_topc(X24))|~r2_hidden(k5_setfam_1(u1_struct_0(X24),esk4_1(X24)),u1_pre_topc(X24))|~r2_hidden(u1_struct_0(X24),u1_pre_topc(X24))|v2_pre_topc(X24)|~l1_pre_topc(X24)))&(~r2_hidden(k9_subset_1(u1_struct_0(X24),esk5_1(X24),esk6_1(X24)),u1_pre_topc(X24))|~r2_hidden(k5_setfam_1(u1_struct_0(X24),esk4_1(X24)),u1_pre_topc(X24))|~r2_hidden(u1_struct_0(X24),u1_pre_topc(X24))|v2_pre_topc(X24)|~l1_pre_topc(X24)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])).
fof(c_0_52, plain, ![X31]:(~v2_pre_topc(X31)|~l1_pre_topc(X31)|r2_hidden(k1_xboole_0,u1_pre_topc(X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_pre_topc])])).
cnf(c_0_53, negated_conjecture, (~v4_pre_topc(X1,esk1_0)|~v3_pre_topc(X1,esk1_0)|~r2_hidden(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_54, plain, (v4_pre_topc(u1_struct_0(X1),X1)|~v3_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_35])])).
cnf(c_0_55, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])])).
cnf(c_0_56, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_49, c_0_50])).
fof(c_0_57, plain, ![X32, X33]:((~v3_pre_topc(X33,X32)|r2_hidden(X33,u1_pre_topc(X32))|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32)))|~l1_pre_topc(X32))&(~r2_hidden(X33,u1_pre_topc(X32))|v3_pre_topc(X33,X32)|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32)))|~l1_pre_topc(X32))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_pre_topc])])])])).
cnf(c_0_58, plain, (r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_59, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_60, plain, (r2_hidden(k1_xboole_0,u1_pre_topc(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_61, negated_conjecture, (~v3_pre_topc(u1_struct_0(esk1_0),esk1_0)|~v3_pre_topc(esk3_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55]), c_0_56]), c_0_48])])).
cnf(c_0_62, plain, (v3_pre_topc(X1,X2)|~r2_hidden(X1,u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_63, negated_conjecture, (r2_hidden(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_48])])).
cnf(c_0_64, plain, (r2_hidden(esk3_0,u1_pre_topc(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[c_0_60, c_0_23])).
cnf(c_0_65, negated_conjecture, (~v3_pre_topc(esk3_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63]), c_0_56]), c_0_48])])).
cnf(c_0_66, negated_conjecture, (r2_hidden(esk3_0,u1_pre_topc(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_59]), c_0_48])])).
cnf(c_0_67, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_62]), c_0_66]), c_0_35]), c_0_48])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 68
# Proof object clause steps            : 37
# Proof object formula steps           : 31
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 29
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 43
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 42
# Processed clauses                    : 82
# ...of these trivial                  : 1
# ...subsumed                          : 3
# ...remaining for further processing  : 78
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 19
# ...of the previous two non-trivial   : 15
# Contextual simplify-reflections      : 0
# Paramodulations                      : 19
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 35
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 20
# Current number of unprocessed clauses: 16
# ...number of literals in the above   : 83
# Current number of archived formulas  : 0
# Current number of archived clauses   : 44
# Clause-clause subsumption calls (NU) : 598
# Rec. Clause-clause subsumption calls : 105
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 8
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4002

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28608 pages
