# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t21_xboole_1, axiom, ![X1, X2]:k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/4a0c03a1be71479981e293a1f7624cd5.E_input', t21_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/4a0c03a1be71479981e293a1f7624cd5.E_input', t48_xboole_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/4a0c03a1be71479981e293a1f7624cd5.E_input', t12_xboole_1)).
fof(t148_zfmisc_1, conjecture, ![X1, X2, X3, X4]:((r1_tarski(X1,X2)&(k3_xboole_0(X2,X3)=k1_tarski(X4)&r2_hidden(X4,X1)))=>k3_xboole_0(X1,X3)=k1_tarski(X4)), file('data/MPTP/proofs/4a0c03a1be71479981e293a1f7624cd5.E_input', t148_zfmisc_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/4a0c03a1be71479981e293a1f7624cd5.E_input', d5_xboole_0)).
fof(t47_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/4a0c03a1be71479981e293a1f7624cd5.E_input', t47_xboole_1)).
fof(t142_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r1_tarski(X4,k1_enumset1(X1,X2,X3))<=>~((~(X4=k1_xboole_0)&(~(X4=k1_tarski(X1))&(~(X4=k1_tarski(X2))&(~(X4=k1_tarski(X3))&(~(X4=k2_tarski(X1,X2))&(~(X4=k2_tarski(X2,X3))&(~(X4=k2_tarski(X1,X3))&~(X4=k1_enumset1(X1,X2,X3))))))))))), file('data/MPTP/proofs/4a0c03a1be71479981e293a1f7624cd5.E_input', t142_zfmisc_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/4a0c03a1be71479981e293a1f7624cd5.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/4a0c03a1be71479981e293a1f7624cd5.E_input', t72_enumset1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/4a0c03a1be71479981e293a1f7624cd5.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/4a0c03a1be71479981e293a1f7624cd5.E_input', t70_enumset1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/4a0c03a1be71479981e293a1f7624cd5.E_input', idempotence_k3_xboole_0)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/4a0c03a1be71479981e293a1f7624cd5.E_input', d10_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/4a0c03a1be71479981e293a1f7624cd5.E_input', d3_tarski)).
fof(antisymmetry_r2_hidden, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>~(r2_hidden(X2,X1))), file('data/MPTP/proofs/4a0c03a1be71479981e293a1f7624cd5.E_input', antisymmetry_r2_hidden)).
fof(t65_zfmisc_1, axiom, ![X1, X2]:(k4_xboole_0(X1,k1_tarski(X2))=X1<=>~(r2_hidden(X2,X1))), file('data/MPTP/proofs/4a0c03a1be71479981e293a1f7624cd5.E_input', t65_zfmisc_1)).
fof(t16_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3)), file('data/MPTP/proofs/4a0c03a1be71479981e293a1f7624cd5.E_input', t16_xboole_1)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/4a0c03a1be71479981e293a1f7624cd5.E_input', t17_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/4a0c03a1be71479981e293a1f7624cd5.E_input', commutativity_k3_xboole_0)).
fof(c_0_19, plain, ![X5, X6]:k3_xboole_0(X5,k2_xboole_0(X5,X6))=X5, inference(variable_rename,[status(thm)],[t21_xboole_1])).
fof(c_0_20, plain, ![X7, X8]:k4_xboole_0(X7,k4_xboole_0(X7,X8))=k3_xboole_0(X7,X8), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_21, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_22, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_23, plain, ![X9, X10]:(~r1_tarski(X9,X10)|k2_xboole_0(X9,X10)=X10), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
fof(c_0_24, negated_conjecture, ~(![X1, X2, X3, X4]:((r1_tarski(X1,X2)&(k3_xboole_0(X2,X3)=k1_tarski(X4)&r2_hidden(X4,X1)))=>k3_xboole_0(X1,X3)=k1_tarski(X4))), inference(assume_negation,[status(cth)],[t148_zfmisc_1])).
fof(c_0_25, plain, ![X15, X16, X17, X18, X19, X20, X21, X22]:((((r2_hidden(X18,X15)|~r2_hidden(X18,X17)|X17!=k4_xboole_0(X15,X16))&(~r2_hidden(X18,X16)|~r2_hidden(X18,X17)|X17!=k4_xboole_0(X15,X16)))&(~r2_hidden(X19,X15)|r2_hidden(X19,X16)|r2_hidden(X19,X17)|X17!=k4_xboole_0(X15,X16)))&((~r2_hidden(esk5_3(X20,X21,X22),X22)|(~r2_hidden(esk5_3(X20,X21,X22),X20)|r2_hidden(esk5_3(X20,X21,X22),X21))|X22=k4_xboole_0(X20,X21))&((r2_hidden(esk5_3(X20,X21,X22),X20)|r2_hidden(esk5_3(X20,X21,X22),X22)|X22=k4_xboole_0(X20,X21))&(~r2_hidden(esk5_3(X20,X21,X22),X21)|r2_hidden(esk5_3(X20,X21,X22),X22)|X22=k4_xboole_0(X20,X21))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_26, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k2_xboole_0(X1,X2)))=X1), inference(rw,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_28, negated_conjecture, ((r1_tarski(esk1_0,esk2_0)&(k3_xboole_0(esk2_0,esk3_0)=k1_tarski(esk4_0)&r2_hidden(esk4_0,esk1_0)))&k3_xboole_0(esk1_0,esk3_0)!=k1_tarski(esk4_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).
fof(c_0_29, plain, ![X24, X25]:k4_xboole_0(X24,k3_xboole_0(X24,X25))=k4_xboole_0(X24,X25), inference(variable_rename,[status(thm)],[t47_xboole_1])).
cnf(c_0_30, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_32, plain, ![X29, X30, X31, X32]:((~r1_tarski(X32,k1_enumset1(X29,X30,X31))|(X32=k1_xboole_0|(X32=k1_tarski(X29)|(X32=k1_tarski(X30)|(X32=k1_tarski(X31)|(X32=k2_tarski(X29,X30)|(X32=k2_tarski(X30,X31)|(X32=k2_tarski(X29,X31)|X32=k1_enumset1(X29,X30,X31)))))))))&((X32!=k1_xboole_0|r1_tarski(X32,k1_enumset1(X29,X30,X31)))&((X32!=k1_tarski(X29)|r1_tarski(X32,k1_enumset1(X29,X30,X31)))&((X32!=k1_tarski(X30)|r1_tarski(X32,k1_enumset1(X29,X30,X31)))&((X32!=k1_tarski(X31)|r1_tarski(X32,k1_enumset1(X29,X30,X31)))&((X32!=k2_tarski(X29,X30)|r1_tarski(X32,k1_enumset1(X29,X30,X31)))&((X32!=k2_tarski(X30,X31)|r1_tarski(X32,k1_enumset1(X29,X30,X31)))&((X32!=k2_tarski(X29,X31)|r1_tarski(X32,k1_enumset1(X29,X30,X31)))&(X32!=k1_enumset1(X29,X30,X31)|r1_tarski(X32,k1_enumset1(X29,X30,X31))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t142_zfmisc_1])])])])).
fof(c_0_33, plain, ![X33, X34, X35]:k2_enumset1(X33,X33,X34,X35)=k1_enumset1(X33,X34,X35), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_34, plain, ![X36, X37, X38, X39]:k3_enumset1(X36,X36,X37,X38,X39)=k2_enumset1(X36,X37,X38,X39), inference(variable_rename,[status(thm)],[t72_enumset1])).
cnf(c_0_35, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=X1|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_36, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_30])).
cnf(c_0_39, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_31])).
cnf(c_0_40, plain, (r1_tarski(X1,k1_enumset1(X2,X3,X4))|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_41, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_43, plain, ![X48]:k2_tarski(X48,X48)=k1_tarski(X48), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_44, plain, ![X49, X50]:k1_enumset1(X49,X49,X50)=k2_tarski(X49,X50), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_45, negated_conjecture, (k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0))=esk1_0), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_46, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X1,X2)))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_37, c_0_22])).
cnf(c_0_47, plain, (~r2_hidden(X1,k4_xboole_0(X2,k2_xboole_0(X2,X3)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_26]), c_0_39])).
fof(c_0_48, plain, ![X26]:k3_xboole_0(X26,X26)=X26, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
fof(c_0_49, plain, ![X27, X28]:(((r1_tarski(X27,X28)|X27!=X28)&(r1_tarski(X28,X27)|X27!=X28))&(~r1_tarski(X27,X28)|~r1_tarski(X28,X27)|X27=X28)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_50, plain, (r1_tarski(X1,k3_enumset1(X2,X2,X2,X3,X4))|X1!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_41]), c_0_42])).
cnf(c_0_51, negated_conjecture, (k3_xboole_0(esk2_0,esk3_0)=k1_tarski(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_52, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_53, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_54, negated_conjecture, (~r2_hidden(X1,k4_xboole_0(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_45]), c_0_39])).
cnf(c_0_55, negated_conjecture, (k4_xboole_0(esk1_0,esk2_0)=k4_xboole_0(esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_46, c_0_45])).
fof(c_0_56, plain, ![X40, X41, X42, X43, X44]:((~r1_tarski(X40,X41)|(~r2_hidden(X42,X40)|r2_hidden(X42,X41)))&((r2_hidden(esk6_2(X43,X44),X43)|r1_tarski(X43,X44))&(~r2_hidden(esk6_2(X43,X44),X44)|r1_tarski(X43,X44)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_57, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_47, c_0_27])).
cnf(c_0_58, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_60, plain, (r1_tarski(k1_xboole_0,k3_enumset1(X1,X1,X1,X2,X3))), inference(er,[status(thm)],[c_0_50])).
cnf(c_0_61, negated_conjecture, (k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0))=k3_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_52]), c_0_53]), c_0_41]), c_0_22]), c_0_42])).
cnf(c_0_62, negated_conjecture, (~r2_hidden(X1,k4_xboole_0(esk1_0,esk1_0))), inference(rw,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_63, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_64, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r1_tarski(X2,k4_xboole_0(X2,k4_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_57, c_0_46])).
cnf(c_0_65, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X1))=X1), inference(rw,[status(thm)],[c_0_58, c_0_22])).
cnf(c_0_66, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_59])).
fof(c_0_67, plain, ![X51, X52]:(~r2_hidden(X51,X52)|~r2_hidden(X52,X51)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[antisymmetry_r2_hidden])])])).
cnf(c_0_68, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_69, negated_conjecture, (r1_tarski(k1_xboole_0,k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_60, c_0_61])).
fof(c_0_70, plain, ![X46, X47]:((k4_xboole_0(X46,k1_tarski(X47))!=X46|~r2_hidden(X47,X46))&(r2_hidden(X47,X46)|k4_xboole_0(X46,k1_tarski(X47))=X46)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t65_zfmisc_1])])])).
cnf(c_0_71, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_72, negated_conjecture, (r1_tarski(k4_xboole_0(esk1_0,esk1_0),X1)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_73, plain, (~r2_hidden(X1,k4_xboole_0(X2,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_66])])).
cnf(c_0_74, plain, (~r2_hidden(X1,X2)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_75, negated_conjecture, (r2_hidden(X1,k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0)))|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_76, plain, (r2_hidden(X1,X2)|k4_xboole_0(X2,k1_tarski(X1))=X2), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_77, negated_conjecture, (X1=k4_xboole_0(esk1_0,esk1_0)|~r1_tarski(X1,k4_xboole_0(esk1_0,esk1_0))), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_78, plain, (r1_tarski(k4_xboole_0(X1,X1),X2)), inference(spm,[status(thm)],[c_0_73, c_0_63])).
cnf(c_0_79, negated_conjecture, (~r2_hidden(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0)),X1)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_80, plain, (k4_xboole_0(k1_xboole_0,k4_xboole_0(k1_xboole_0,k3_enumset1(X1,X1,X1,X2,X3)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_35, c_0_60])).
cnf(c_0_81, plain, (k4_xboole_0(X2,k3_enumset1(X1,X1,X1,X1,X1))=X2|r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_52]), c_0_53]), c_0_41]), c_0_42])).
cnf(c_0_82, negated_conjecture, (k4_xboole_0(X1,X1)=k4_xboole_0(esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_83, plain, (r1_tarski(X1,k1_enumset1(X2,X4,X3))|X1!=k2_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_84, plain, ![X55, X56, X57]:k3_xboole_0(k3_xboole_0(X55,X56),X57)=k3_xboole_0(X55,k3_xboole_0(X56,X57)), inference(variable_rename,[status(thm)],[t16_xboole_1])).
cnf(c_0_85, negated_conjecture, (~r2_hidden(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0)),k1_xboole_0)), inference(spm,[status(thm)],[c_0_79, c_0_75])).
cnf(c_0_86, plain, (k4_xboole_0(esk1_0,esk1_0)=k1_xboole_0|r2_hidden(X1,k1_xboole_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_82])).
cnf(c_0_87, plain, (r1_tarski(X1,k3_enumset1(X2,X2,X2,X4,X3))|X1!=k3_enumset1(X2,X2,X2,X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_83, c_0_53]), c_0_41]), c_0_41]), c_0_42]), c_0_42])).
cnf(c_0_88, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_89, negated_conjecture, (k4_xboole_0(X1,k4_xboole_0(esk1_0,esk1_0))=X1), inference(spm,[status(thm)],[c_0_65, c_0_82])).
cnf(c_0_90, negated_conjecture, (k4_xboole_0(X1,X1)=k4_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_82, c_0_82])).
cnf(c_0_91, negated_conjecture, (k4_xboole_0(esk1_0,esk1_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_92, plain, (r1_tarski(k3_enumset1(X1,X1,X1,X1,X2),k3_enumset1(X1,X1,X1,X3,X2))), inference(er,[status(thm)],[c_0_87])).
cnf(c_0_93, plain, (k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X3))=k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_88, c_0_22]), c_0_22]), c_0_22]), c_0_22])).
cnf(c_0_94, negated_conjecture, (k4_xboole_0(X1,k4_xboole_0(X2,X2))=X1), inference(spm,[status(thm)],[c_0_89, c_0_82])).
cnf(c_0_95, negated_conjecture, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_96, negated_conjecture, (r1_tarski(k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0)),k3_enumset1(esk4_0,esk4_0,esk4_0,X1,esk4_0))), inference(spm,[status(thm)],[c_0_92, c_0_61])).
cnf(c_0_97, plain, (k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))))=k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X3)), inference(spm,[status(thm)],[c_0_46, c_0_93])).
cnf(c_0_98, negated_conjecture, (k4_xboole_0(esk1_0,esk2_0)=k1_xboole_0), inference(rw,[status(thm)],[c_0_55, c_0_91])).
cnf(c_0_99, negated_conjecture, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_100, negated_conjecture, (k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,k3_enumset1(esk4_0,esk4_0,esk4_0,X1,esk4_0)))))=k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_96]), c_0_93])).
cnf(c_0_101, negated_conjecture, (k4_xboole_0(esk1_0,k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,X1)))=k4_xboole_0(esk1_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_99]), c_0_46]), c_0_99])).
cnf(c_0_102, negated_conjecture, (k4_xboole_0(esk2_0,k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,k3_enumset1(esk4_0,esk4_0,esk4_0,X1,esk4_0))))=k4_xboole_0(esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_100]), c_0_46])).
cnf(c_0_103, plain, (~r2_hidden(X1,k4_xboole_0(k4_xboole_0(X2,k4_xboole_0(X2,X3)),X4))|~r1_tarski(k4_xboole_0(X2,k4_xboole_0(X2,X3)),k4_xboole_0(X2,k4_xboole_0(X2,k4_xboole_0(X3,k4_xboole_0(X3,X4)))))), inference(spm,[status(thm)],[c_0_64, c_0_93])).
cnf(c_0_104, negated_conjecture, (k4_xboole_0(esk1_0,k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,k3_enumset1(esk4_0,esk4_0,esk4_0,X1,esk4_0))))=k4_xboole_0(esk1_0,esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_101])).
cnf(c_0_105, negated_conjecture, (~r2_hidden(X1,k4_xboole_0(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0)),k3_enumset1(esk4_0,esk4_0,esk4_0,X2,esk4_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_66])])).
cnf(c_0_106, plain, (k4_xboole_0(X1,k1_tarski(X2))!=X1|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
fof(c_0_107, plain, ![X58, X59]:r1_tarski(k3_xboole_0(X58,X59),X58), inference(variable_rename,[status(thm)],[t17_xboole_1])).
cnf(c_0_108, negated_conjecture, (~r2_hidden(X1,k4_xboole_0(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0)),k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0))))), inference(spm,[status(thm)],[c_0_105, c_0_61])).
cnf(c_0_109, plain, (k4_xboole_0(X1,k3_enumset1(X2,X2,X2,X2,X2))!=X1|~r2_hidden(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_106, c_0_52]), c_0_53]), c_0_41]), c_0_42])).
cnf(c_0_110, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_107])).
fof(c_0_111, plain, ![X53, X54]:k3_xboole_0(X53,X54)=k3_xboole_0(X54,X53), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_112, negated_conjecture, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_91]), c_0_91])).
cnf(c_0_113, negated_conjecture, (r1_tarski(k4_xboole_0(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0)),k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0))),X1)), inference(spm,[status(thm)],[c_0_108, c_0_63])).
cnf(c_0_114, negated_conjecture, (k4_xboole_0(X1,k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0)))!=X1|~r2_hidden(esk4_0,X1)), inference(spm,[status(thm)],[c_0_109, c_0_61])).
cnf(c_0_115, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X1)), inference(rw,[status(thm)],[c_0_110, c_0_22])).
cnf(c_0_116, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_117, negated_conjecture, (k4_xboole_0(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0)),k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_112, c_0_113])).
cnf(c_0_118, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_119, negated_conjecture, (k4_xboole_0(esk2_0,esk3_0)!=esk2_0|~r2_hidden(esk4_0,esk2_0)), inference(spm,[status(thm)],[c_0_114, c_0_46])).
cnf(c_0_120, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_68, c_0_36])).
cnf(c_0_121, negated_conjecture, (r2_hidden(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_122, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|X1=k1_tarski(X3)|X1=k1_tarski(X4)|X1=k2_tarski(X2,X3)|X1=k2_tarski(X3,X4)|X1=k2_tarski(X2,X4)|X1=k1_enumset1(X2,X3,X4)|~r1_tarski(X1,k1_enumset1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_123, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_115, c_0_46])).
cnf(c_0_124, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_116, c_0_22]), c_0_22])).
cnf(c_0_125, negated_conjecture, (k3_xboole_0(esk1_0,esk3_0)!=k1_tarski(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_126, negated_conjecture, (k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))!=k1_xboole_0|~r2_hidden(esk4_0,k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0)))), inference(spm,[status(thm)],[c_0_114, c_0_117])).
cnf(c_0_127, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_118])).
cnf(c_0_128, negated_conjecture, (k4_xboole_0(esk2_0,esk3_0)!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120]), c_0_121])])).
cnf(c_0_129, plain, (X1=k1_xboole_0|X1=k3_enumset1(X4,X4,X4,X4,X4)|X1=k3_enumset1(X3,X3,X3,X3,X4)|X1=k3_enumset1(X3,X3,X3,X3,X3)|X1=k3_enumset1(X2,X2,X2,X3,X4)|X1=k3_enumset1(X2,X2,X2,X2,X4)|X1=k3_enumset1(X2,X2,X2,X2,X3)|X1=k3_enumset1(X2,X2,X2,X2,X2)|~r1_tarski(X1,k3_enumset1(X2,X2,X2,X3,X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_122, c_0_52]), c_0_52]), c_0_52]), c_0_53]), c_0_53]), c_0_53]), c_0_53]), c_0_53]), c_0_53]), c_0_41]), c_0_41]), c_0_41]), c_0_41]), c_0_41]), c_0_41]), c_0_41]), c_0_41]), c_0_42]), c_0_42]), c_0_42]), c_0_42]), c_0_42]), c_0_42]), c_0_42]), c_0_42])).
cnf(c_0_130, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_123, c_0_124])).
cnf(c_0_131, negated_conjecture, (k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))!=k3_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_125, c_0_52]), c_0_53]), c_0_41]), c_0_22]), c_0_42])).
cnf(c_0_132, negated_conjecture, (r2_hidden(esk4_0,k4_xboole_0(esk1_0,esk3_0))|k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_127]), c_0_121])])).
cnf(c_0_133, negated_conjecture, (r2_hidden(esk4_0,esk3_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_81]), c_0_94]), c_0_128])).
cnf(c_0_134, negated_conjecture, (X1=k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0))|X1=k1_xboole_0|~r1_tarski(X1,k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_61]), c_0_61])])).
cnf(c_0_135, negated_conjecture, (r1_tarski(k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0)),k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_117]), c_0_99])).
cnf(c_0_136, negated_conjecture, (k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk3_0))!=k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))), inference(rw,[status(thm)],[c_0_131, c_0_61])).
cnf(c_0_137, negated_conjecture, (k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk3_0))!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_132]), c_0_133])])).
cnf(c_0_138, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_136]), c_0_137]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 139
# Proof object clause steps            : 100
# Proof object formula steps           : 39
# Proof object conjectures             : 49
# Proof object clause conjectures      : 46
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 29
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 47
# Proof object simplifying inferences  : 98
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 35
# Processed clauses                    : 1579
# ...of these trivial                  : 39
# ...subsumed                          : 1095
# ...remaining for further processing  : 445
# Other redundant clauses eliminated   : 13
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 21
# Backward-rewritten                   : 90
# Generated clauses                    : 8596
# ...of the previous two non-trivial   : 6839
# Contextual simplify-reflections      : 4
# Paramodulations                      : 8576
# Factorizations                       : 6
# Equation resolutions                 : 13
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 287
#    Positive orientable unit clauses  : 61
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 38
#    Non-unit-clauses                  : 185
# Current number of unprocessed clauses: 5131
# ...number of literals in the above   : 10438
# Current number of archived formulas  : 0
# Current number of archived clauses   : 150
# Clause-clause subsumption calls (NU) : 6443
# Rec. Clause-clause subsumption calls : 5428
# Non-unit clause-clause subsumptions  : 436
# Unit Clause-clause subsumption calls : 739
# Rewrite failures with RHS unbound    : 37
# BW rewrite match attempts            : 559
# BW rewrite match successes           : 129
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 139145

# -------------------------------------------------
# User time                : 0.073 s
# System time              : 0.007 s
# Total time               : 0.080 s
# Maximum resident set size: 28592 pages
