# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/4a18fb59e0e34b2eba313a10df2e992c.E_input', d5_relat_1)).
fof(d8_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(X3=k5_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>?[X6]:(r2_hidden(k4_tarski(X4,X6),X1)&r2_hidden(k4_tarski(X6,X5),X2))))))), file('data/MPTP/proofs/4a18fb59e0e34b2eba313a10df2e992c.E_input', d8_relat_1)).
fof(dt_k5_relat_1, axiom, ![X1, X2]:((v1_relat_1(X1)&v1_relat_1(X2))=>v1_relat_1(k5_relat_1(X1,X2))), file('data/MPTP/proofs/4a18fb59e0e34b2eba313a10df2e992c.E_input', dt_k5_relat_1)).
fof(t25_funct_1, conjecture, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k2_relat_1(k5_relat_1(X3,X2)))=>r2_hidden(X1,k2_relat_1(X2))))), file('data/MPTP/proofs/4a18fb59e0e34b2eba313a10df2e992c.E_input', t25_funct_1)).
fof(c_0_4, plain, ![X10, X11, X12, X14, X15, X16, X17, X19]:(((~r2_hidden(X12,X11)|r2_hidden(k4_tarski(esk4_3(X10,X11,X12),X12),X10)|X11!=k2_relat_1(X10))&(~r2_hidden(k4_tarski(X15,X14),X10)|r2_hidden(X14,X11)|X11!=k2_relat_1(X10)))&((~r2_hidden(esk5_2(X16,X17),X17)|~r2_hidden(k4_tarski(X19,esk5_2(X16,X17)),X16)|X17=k2_relat_1(X16))&(r2_hidden(esk5_2(X16,X17),X17)|r2_hidden(k4_tarski(esk6_2(X16,X17),esk5_2(X16,X17)),X16)|X17=k2_relat_1(X16)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
fof(c_0_5, plain, ![X21, X22, X23, X24, X25, X27, X28, X29, X32]:((((r2_hidden(k4_tarski(X24,esk7_5(X21,X22,X23,X24,X25)),X21)|~r2_hidden(k4_tarski(X24,X25),X23)|X23!=k5_relat_1(X21,X22)|~v1_relat_1(X23)|~v1_relat_1(X22)|~v1_relat_1(X21))&(r2_hidden(k4_tarski(esk7_5(X21,X22,X23,X24,X25),X25),X22)|~r2_hidden(k4_tarski(X24,X25),X23)|X23!=k5_relat_1(X21,X22)|~v1_relat_1(X23)|~v1_relat_1(X22)|~v1_relat_1(X21)))&(~r2_hidden(k4_tarski(X27,X29),X21)|~r2_hidden(k4_tarski(X29,X28),X22)|r2_hidden(k4_tarski(X27,X28),X23)|X23!=k5_relat_1(X21,X22)|~v1_relat_1(X23)|~v1_relat_1(X22)|~v1_relat_1(X21)))&((~r2_hidden(k4_tarski(esk8_3(X21,X22,X23),esk9_3(X21,X22,X23)),X23)|(~r2_hidden(k4_tarski(esk8_3(X21,X22,X23),X32),X21)|~r2_hidden(k4_tarski(X32,esk9_3(X21,X22,X23)),X22))|X23=k5_relat_1(X21,X22)|~v1_relat_1(X23)|~v1_relat_1(X22)|~v1_relat_1(X21))&((r2_hidden(k4_tarski(esk8_3(X21,X22,X23),esk10_3(X21,X22,X23)),X21)|r2_hidden(k4_tarski(esk8_3(X21,X22,X23),esk9_3(X21,X22,X23)),X23)|X23=k5_relat_1(X21,X22)|~v1_relat_1(X23)|~v1_relat_1(X22)|~v1_relat_1(X21))&(r2_hidden(k4_tarski(esk10_3(X21,X22,X23),esk9_3(X21,X22,X23)),X22)|r2_hidden(k4_tarski(esk8_3(X21,X22,X23),esk9_3(X21,X22,X23)),X23)|X23=k5_relat_1(X21,X22)|~v1_relat_1(X23)|~v1_relat_1(X22)|~v1_relat_1(X21))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_relat_1])])])])])])).
fof(c_0_6, plain, ![X34, X35]:(~v1_relat_1(X34)|~v1_relat_1(X35)|v1_relat_1(k5_relat_1(X34,X35))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_relat_1])])).
cnf(c_0_7, plain, (r2_hidden(X2,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, plain, (r2_hidden(k4_tarski(esk7_5(X1,X2,X3,X4,X5),X5),X2)|~r2_hidden(k4_tarski(X4,X5),X3)|X3!=k5_relat_1(X1,X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (v1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)), inference(er,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (r2_hidden(k4_tarski(esk7_5(X1,X2,k5_relat_1(X1,X2),X3,X4),X4),X2)|~r2_hidden(k4_tarski(X3,X4),k5_relat_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_8]), c_0_9])).
cnf(c_0_12, plain, (r2_hidden(k4_tarski(esk4_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_13, negated_conjecture, ~(![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k2_relat_1(k5_relat_1(X3,X2)))=>r2_hidden(X1,k2_relat_1(X2)))))), inference(assume_negation,[status(cth)],[t25_funct_1])).
cnf(c_0_14, plain, (r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),k5_relat_1(X4,X2))|~v1_relat_1(X2)|~v1_relat_1(X4)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15, plain, (r2_hidden(k4_tarski(esk4_3(X1,k2_relat_1(X1),X2),X2),X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_12])).
fof(c_0_16, negated_conjecture, ((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&((v1_relat_1(esk3_0)&v1_funct_1(esk3_0))&(r2_hidden(esk1_0,k2_relat_1(k5_relat_1(esk3_0,esk2_0)))&~r2_hidden(esk1_0,k2_relat_1(esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
cnf(c_0_17, plain, (r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(X1,k2_relat_1(k5_relat_1(X3,X2)))|~v1_relat_1(X2)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_18, negated_conjecture, (r2_hidden(esk1_0,k2_relat_1(k5_relat_1(esk3_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_19, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, negated_conjecture, (~r2_hidden(esk1_0,k2_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19]), c_0_20])]), c_0_21]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 23
# Proof object clause steps            : 14
# Proof object formula steps           : 9
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 45
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 45
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 39
# ...of the previous two non-trivial   : 29
# Contextual simplify-reflections      : 3
# Paramodulations                      : 32
# Factorizations                       : 2
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 23
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 17
# Current number of unprocessed clauses: 15
# ...number of literals in the above   : 85
# Current number of archived formulas  : 0
# Current number of archived clauses   : 17
# Clause-clause subsumption calls (NU) : 225
# Rec. Clause-clause subsumption calls : 26
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2416

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.003 s
# Total time               : 0.010 s
# Maximum resident set size: 28788 pages
