# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_funct_2, axiom, ![X1, X2, X3]:(X3=k1_funct_2(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:((v1_relat_1(X5)&v1_funct_1(X5))&(X4=X5&(k1_relat_1(X5)=X1&r1_tarski(k2_relat_1(X5),X2)))))), file('data/MPTP/proofs/4a1c8e2a147944d19c87aaa2f79b9b3a.E_input', d2_funct_2)).
fof(t121_funct_2, conjecture, ![X1, X2, X3]:(r2_hidden(X3,k1_funct_2(X1,X2))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))), file('data/MPTP/proofs/4a1c8e2a147944d19c87aaa2f79b9b3a.E_input', t121_funct_2)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/4a1c8e2a147944d19c87aaa2f79b9b3a.E_input', d5_relat_1)).
fof(t11_relset_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>((r1_tarski(k1_relat_1(X3),X1)&r1_tarski(k2_relat_1(X3),X2))=>m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))), file('data/MPTP/proofs/4a1c8e2a147944d19c87aaa2f79b9b3a.E_input', t11_relset_1)).
fof(d1_funct_1, axiom, ![X1]:(v1_funct_1(X1)<=>![X2, X3, X4]:((r2_hidden(k4_tarski(X2,X3),X1)&r2_hidden(k4_tarski(X2,X4),X1))=>X3=X4)), file('data/MPTP/proofs/4a1c8e2a147944d19c87aaa2f79b9b3a.E_input', d1_funct_1)).
fof(d5_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4))))), file('data/MPTP/proofs/4a1c8e2a147944d19c87aaa2f79b9b3a.E_input', d5_funct_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/4a1c8e2a147944d19c87aaa2f79b9b3a.E_input', d1_xboole_0)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/4a1c8e2a147944d19c87aaa2f79b9b3a.E_input', d10_xboole_0)).
fof(t3_funct_2, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_funct_1(X1)&(v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))))), file('data/MPTP/proofs/4a1c8e2a147944d19c87aaa2f79b9b3a.E_input', t3_funct_2)).
fof(cc1_partfun1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_partfun1(X3,X1))), file('data/MPTP/proofs/4a1c8e2a147944d19c87aaa2f79b9b3a.E_input', cc1_partfun1)).
fof(cc1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_partfun1(X3,X1))=>(v1_funct_1(X3)&v1_funct_2(X3,X1,X2)))), file('data/MPTP/proofs/4a1c8e2a147944d19c87aaa2f79b9b3a.E_input', cc1_funct_2)).
fof(cc5_funct_2, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_funct_2(X3,X1,X2))=>(v1_funct_1(X3)&v1_partfun1(X3,X1))))), file('data/MPTP/proofs/4a1c8e2a147944d19c87aaa2f79b9b3a.E_input', cc5_funct_2)).
fof(c_0_12, plain, ![X51, X52, X53, X54, X56, X57, X58, X59, X60, X62]:(((((v1_relat_1(esk15_4(X51,X52,X53,X54))|~r2_hidden(X54,X53)|X53!=k1_funct_2(X51,X52))&(v1_funct_1(esk15_4(X51,X52,X53,X54))|~r2_hidden(X54,X53)|X53!=k1_funct_2(X51,X52)))&((X54=esk15_4(X51,X52,X53,X54)|~r2_hidden(X54,X53)|X53!=k1_funct_2(X51,X52))&((k1_relat_1(esk15_4(X51,X52,X53,X54))=X51|~r2_hidden(X54,X53)|X53!=k1_funct_2(X51,X52))&(r1_tarski(k2_relat_1(esk15_4(X51,X52,X53,X54)),X52)|~r2_hidden(X54,X53)|X53!=k1_funct_2(X51,X52)))))&(~v1_relat_1(X57)|~v1_funct_1(X57)|(X56!=X57|(k1_relat_1(X57)!=X51|~r1_tarski(k2_relat_1(X57),X52)))|r2_hidden(X56,X53)|X53!=k1_funct_2(X51,X52)))&((~r2_hidden(esk16_3(X58,X59,X60),X60)|(~v1_relat_1(X62)|~v1_funct_1(X62)|(esk16_3(X58,X59,X60)!=X62|(k1_relat_1(X62)!=X58|~r1_tarski(k2_relat_1(X62),X59))))|X60=k1_funct_2(X58,X59))&(((v1_relat_1(esk17_3(X58,X59,X60))|r2_hidden(esk16_3(X58,X59,X60),X60)|X60=k1_funct_2(X58,X59))&(v1_funct_1(esk17_3(X58,X59,X60))|r2_hidden(esk16_3(X58,X59,X60),X60)|X60=k1_funct_2(X58,X59)))&((esk16_3(X58,X59,X60)=esk17_3(X58,X59,X60)|r2_hidden(esk16_3(X58,X59,X60),X60)|X60=k1_funct_2(X58,X59))&((k1_relat_1(esk17_3(X58,X59,X60))=X58|r2_hidden(esk16_3(X58,X59,X60),X60)|X60=k1_funct_2(X58,X59))&(r1_tarski(k2_relat_1(esk17_3(X58,X59,X60)),X59)|r2_hidden(esk16_3(X58,X59,X60),X60)|X60=k1_funct_2(X58,X59))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_funct_2])])])])])])).
fof(c_0_13, negated_conjecture, ~(![X1, X2, X3]:(r2_hidden(X3,k1_funct_2(X1,X2))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))), inference(assume_negation,[status(cth)],[t121_funct_2])).
cnf(c_0_14, plain, (v1_funct_1(esk15_4(X1,X2,X3,X4))|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15, plain, (X1=esk15_4(X2,X3,X4,X1)|~r2_hidden(X1,X4)|X4!=k1_funct_2(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_16, plain, ![X9, X10, X11, X13, X14, X15, X16, X18]:(((~r2_hidden(X11,X10)|r2_hidden(k4_tarski(esk4_3(X9,X10,X11),X11),X9)|X10!=k2_relat_1(X9))&(~r2_hidden(k4_tarski(X14,X13),X9)|r2_hidden(X13,X10)|X10!=k2_relat_1(X9)))&((~r2_hidden(esk5_2(X15,X16),X16)|~r2_hidden(k4_tarski(X18,esk5_2(X15,X16)),X15)|X16=k2_relat_1(X15))&(r2_hidden(esk5_2(X15,X16),X16)|r2_hidden(k4_tarski(esk6_2(X15,X16),esk5_2(X15,X16)),X15)|X16=k2_relat_1(X15)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
fof(c_0_17, negated_conjecture, (r2_hidden(esk3_0,k1_funct_2(esk1_0,esk2_0))&(~v1_funct_1(esk3_0)|(~v1_funct_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
fof(c_0_18, plain, ![X20, X21, X22]:(~v1_relat_1(X22)|(~r1_tarski(k1_relat_1(X22),X20)|~r1_tarski(k2_relat_1(X22),X21)|m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(X20,X21))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_relset_1])])).
cnf(c_0_19, plain, (v1_funct_1(esk15_4(X1,X2,k1_funct_2(X1,X2),X3))|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (esk15_4(X1,X2,k1_funct_2(X1,X2),X3)=X3|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (v1_relat_1(esk15_4(X1,X2,X3,X4))|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22, plain, (k1_relat_1(esk15_4(X1,X2,X3,X4))=X1|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23, plain, (r2_hidden(X2,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_24, plain, ![X26, X27, X28, X29, X30]:((~v1_funct_1(X26)|(~r2_hidden(k4_tarski(X27,X28),X26)|~r2_hidden(k4_tarski(X27,X29),X26)|X28=X29))&(((r2_hidden(k4_tarski(esk8_1(X30),esk9_1(X30)),X30)|v1_funct_1(X30))&(r2_hidden(k4_tarski(esk8_1(X30),esk10_1(X30)),X30)|v1_funct_1(X30)))&(esk9_1(X30)!=esk10_1(X30)|v1_funct_1(X30)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_1])])])])])])).
cnf(c_0_25, negated_conjecture, (~v1_funct_1(esk3_0)|~v1_funct_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, plain, (v1_funct_1(X1)|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_28, negated_conjecture, (r2_hidden(esk3_0,k1_funct_2(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_29, plain, (v1_relat_1(esk15_4(X1,X2,k1_funct_2(X1,X2),X3))|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_21])).
cnf(c_0_30, plain, (k1_relat_1(esk15_4(X1,X2,k1_funct_2(X1,X2),X3))=X1|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_22])).
fof(c_0_31, plain, ![X41, X42, X43, X45, X46, X47, X49]:((((r2_hidden(esk12_3(X41,X42,X43),k1_relat_1(X41))|~r2_hidden(X43,X42)|X42!=k2_relat_1(X41)|(~v1_relat_1(X41)|~v1_funct_1(X41)))&(X43=k1_funct_1(X41,esk12_3(X41,X42,X43))|~r2_hidden(X43,X42)|X42!=k2_relat_1(X41)|(~v1_relat_1(X41)|~v1_funct_1(X41))))&(~r2_hidden(X46,k1_relat_1(X41))|X45!=k1_funct_1(X41,X46)|r2_hidden(X45,X42)|X42!=k2_relat_1(X41)|(~v1_relat_1(X41)|~v1_funct_1(X41))))&((~r2_hidden(esk13_2(X41,X47),X47)|(~r2_hidden(X49,k1_relat_1(X41))|esk13_2(X41,X47)!=k1_funct_1(X41,X49))|X47=k2_relat_1(X41)|(~v1_relat_1(X41)|~v1_funct_1(X41)))&((r2_hidden(esk14_2(X41,X47),k1_relat_1(X41))|r2_hidden(esk13_2(X41,X47),X47)|X47=k2_relat_1(X41)|(~v1_relat_1(X41)|~v1_funct_1(X41)))&(esk13_2(X41,X47)=k1_funct_1(X41,esk14_2(X41,X47))|r2_hidden(esk13_2(X41,X47),X47)|X47=k2_relat_1(X41)|(~v1_relat_1(X41)|~v1_funct_1(X41)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_funct_1])])])])])])).
fof(c_0_32, plain, ![X37, X38, X39]:((~v1_xboole_0(X37)|~r2_hidden(X38,X37))&(r2_hidden(esk11_1(X39),X39)|v1_xboole_0(X39))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_33, plain, (r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)), inference(er,[status(thm)],[c_0_23])).
cnf(c_0_34, plain, (r2_hidden(k4_tarski(esk8_1(X1),esk9_1(X1)),X1)|v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_35, negated_conjecture, (~r1_tarski(k2_relat_1(esk3_0),esk2_0)|~r1_tarski(k1_relat_1(esk3_0),esk1_0)|~v1_relat_1(esk3_0)|~v1_funct_2(esk3_0,esk1_0,esk2_0)|~v1_funct_1(esk3_0)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_36, negated_conjecture, (v1_funct_1(esk3_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_37, plain, (v1_relat_1(X1)|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(spm,[status(thm)],[c_0_29, c_0_20])).
fof(c_0_38, plain, ![X64, X65]:(((r1_tarski(X64,X65)|X64!=X65)&(r1_tarski(X65,X64)|X64!=X65))&(~r1_tarski(X64,X65)|~r1_tarski(X65,X64)|X64=X65)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_39, plain, (r1_tarski(k2_relat_1(esk15_4(X1,X2,X3,X4)),X2)|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_40, plain, ![X25]:((v1_funct_1(X25)|(~v1_relat_1(X25)|~v1_funct_1(X25)))&((v1_funct_2(X25,k1_relat_1(X25),k2_relat_1(X25))|(~v1_relat_1(X25)|~v1_funct_1(X25)))&(m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X25),k2_relat_1(X25))))|(~v1_relat_1(X25)|~v1_funct_1(X25))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_funct_2])])])).
cnf(c_0_41, plain, (k1_relat_1(X1)=X2|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(spm,[status(thm)],[c_0_30, c_0_20])).
fof(c_0_42, plain, ![X66, X67, X68]:(~v1_xboole_0(X66)|(~m1_subset_1(X68,k1_zfmisc_1(k2_zfmisc_1(X66,X67)))|v1_partfun1(X68,X66))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_partfun1])])])).
cnf(c_0_43, plain, (r2_hidden(X3,X4)|~r2_hidden(X1,k1_relat_1(X2))|X3!=k1_funct_1(X2,X1)|X4!=k2_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_44, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_45, plain, (v1_funct_1(X1)|r2_hidden(esk9_1(X1),k2_relat_1(X1))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_46, negated_conjecture, (~r1_tarski(k2_relat_1(esk3_0),esk2_0)|~r1_tarski(k1_relat_1(esk3_0),esk1_0)|~v1_relat_1(esk3_0)|~v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_36])])).
cnf(c_0_47, negated_conjecture, (v1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_37, c_0_28])).
cnf(c_0_48, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_49, plain, (r1_tarski(k2_relat_1(esk15_4(X1,X2,k1_funct_2(X1,X2),X3)),X2)|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_39])).
fof(c_0_50, plain, ![X34, X35, X36]:((v1_funct_1(X36)|(~v1_funct_1(X36)|~v1_partfun1(X36,X34))|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(X34,X35))))&(v1_funct_2(X36,X34,X35)|(~v1_funct_1(X36)|~v1_partfun1(X36,X34))|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(X34,X35))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_2])])])).
fof(c_0_51, plain, ![X71, X72, X73]:((v1_funct_1(X73)|(~v1_funct_1(X73)|~v1_funct_2(X73,X71,X72))|~m1_subset_1(X73,k1_zfmisc_1(k2_zfmisc_1(X71,X72)))|v1_xboole_0(X72))&(v1_partfun1(X73,X71)|(~v1_funct_1(X73)|~v1_funct_2(X73,X71,X72))|~m1_subset_1(X73,k1_zfmisc_1(k2_zfmisc_1(X71,X72)))|v1_xboole_0(X72))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_funct_2])])])])])).
cnf(c_0_52, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_53, negated_conjecture, (k1_relat_1(esk3_0)=esk1_0), inference(spm,[status(thm)],[c_0_41, c_0_28])).
cnf(c_0_54, plain, (v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_55, plain, (v1_partfun1(X2,X1)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_56, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_43])])).
cnf(c_0_57, plain, (v1_funct_1(X1)|~v1_xboole_0(k2_relat_1(X1))), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_58, negated_conjecture, (~r1_tarski(k2_relat_1(esk3_0),esk2_0)|~r1_tarski(k1_relat_1(esk3_0),esk1_0)|~v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_47])])).
cnf(c_0_59, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_48])).
cnf(c_0_60, plain, (r1_tarski(k2_relat_1(X1),X2)|~r2_hidden(X1,k1_funct_2(X3,X2))), inference(spm,[status(thm)],[c_0_49, c_0_20])).
cnf(c_0_61, plain, (v1_funct_2(X1,X2,X3)|~v1_funct_1(X1)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_62, plain, (v1_partfun1(X1,X2)|v1_xboole_0(X3)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_63, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_47]), c_0_36])])).
cnf(c_0_64, negated_conjecture, (v1_funct_2(esk3_0,esk1_0,k2_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_53]), c_0_47]), c_0_36])])).
cnf(c_0_65, plain, (v1_partfun1(X1,k1_relat_1(X1))|~v1_xboole_0(k1_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_55, c_0_52])).
cnf(c_0_66, plain, (~v1_xboole_0(k2_relat_1(X1))|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_56]), c_0_57])).
cnf(c_0_67, plain, (r2_hidden(esk11_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_68, negated_conjecture, (~r1_tarski(k2_relat_1(esk3_0),esk2_0)|~v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_53]), c_0_59])])).
cnf(c_0_69, negated_conjecture, (r1_tarski(k2_relat_1(esk3_0),esk2_0)), inference(spm,[status(thm)],[c_0_60, c_0_28])).
cnf(c_0_70, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~r1_tarski(k2_relat_1(X1),X3)|~r1_tarski(k1_relat_1(X1),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_61, c_0_26])).
cnf(c_0_71, negated_conjecture, (v1_xboole_0(k2_relat_1(esk3_0))|v1_partfun1(esk3_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64]), c_0_36])])).
cnf(c_0_72, negated_conjecture, (v1_partfun1(esk3_0,esk1_0)|~v1_xboole_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_53]), c_0_47]), c_0_36])])).
cnf(c_0_73, plain, (v1_xboole_0(k1_relat_1(X1))|~v1_xboole_0(k2_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_74, negated_conjecture, (~v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_69])])).
cnf(c_0_75, negated_conjecture, (v1_xboole_0(k2_relat_1(esk3_0))|v1_funct_2(esk3_0,esk1_0,X1)|~r1_tarski(k2_relat_1(esk3_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_53]), c_0_59]), c_0_47]), c_0_36])])).
cnf(c_0_76, negated_conjecture, (v1_funct_2(esk3_0,esk1_0,X1)|~v1_xboole_0(esk1_0)|~r1_tarski(k2_relat_1(esk3_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_72]), c_0_53]), c_0_59]), c_0_47]), c_0_36])])).
cnf(c_0_77, negated_conjecture, (v1_xboole_0(esk1_0)|~v1_xboole_0(k2_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_53]), c_0_47])])).
cnf(c_0_78, negated_conjecture, (v1_xboole_0(k2_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_69])])).
cnf(c_0_79, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_76]), c_0_69])])).
cnf(c_0_80, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_78])]), c_0_79]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 81
# Proof object clause steps            : 56
# Proof object formula steps           : 25
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 24
# Proof object simplifying inferences  : 50
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 44
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 41
# Processed clauses                    : 389
# ...of these trivial                  : 0
# ...subsumed                          : 146
# ...remaining for further processing  : 243
# Other redundant clauses eliminated   : 17
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 21
# Generated clauses                    : 827
# ...of the previous two non-trivial   : 815
# Contextual simplify-reflections      : 5
# Paramodulations                      : 813
# Factorizations                       : 0
# Equation resolutions                 : 17
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 167
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 151
# Current number of unprocessed clauses: 495
# ...number of literals in the above   : 1809
# Current number of archived formulas  : 0
# Current number of archived clauses   : 62
# Clause-clause subsumption calls (NU) : 5382
# Rec. Clause-clause subsumption calls : 2739
# Non-unit clause-clause subsumptions  : 151
# Unit Clause-clause subsumption calls : 122
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 14827

# -------------------------------------------------
# User time                : 0.021 s
# System time              : 0.000 s
# Total time               : 0.021 s
# Maximum resident set size: 28648 pages
