# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4))))), file('data/MPTP/proofs/4a57d33d3b264ca0aad8af9515b91e03.E_input', d5_funct_1)).
fof(s3_funct_1__e2_24__funct_1, axiom, ![X1, X2]:?[X3]:((v1_relat_1(X3)&v1_funct_1(X3))&(k1_relat_1(X3)=X1&![X4]:(r2_hidden(X4,X1)=>k1_funct_1(X3,X4)=X2))), file('data/MPTP/proofs/4a57d33d3b264ca0aad8af9515b91e03.E_input', s3_funct_1__e2_24__funct_1)).
fof(t15_funct_1, conjecture, ![X1]:(~(X1=k1_xboole_0)=>![X2]:?[X3]:((v1_relat_1(X3)&v1_funct_1(X3))&(k1_relat_1(X3)=X1&k2_relat_1(X3)=k1_tarski(X2)))), file('data/MPTP/proofs/4a57d33d3b264ca0aad8af9515b91e03.E_input', t15_funct_1)).
fof(t41_zfmisc_1, axiom, ![X1, X2]:~((~(X1=k1_tarski(X2))&(~(X1=k1_xboole_0)&![X3]:~((r2_hidden(X3,X1)&~(X3=X2)))))), file('data/MPTP/proofs/4a57d33d3b264ca0aad8af9515b91e03.E_input', t41_zfmisc_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/4a57d33d3b264ca0aad8af9515b91e03.E_input', d1_tarski)).
fof(c_0_5, plain, ![X5, X6, X7, X9, X10, X11, X13]:((((r2_hidden(esk1_3(X5,X6,X7),k1_relat_1(X5))|~r2_hidden(X7,X6)|X6!=k2_relat_1(X5)|(~v1_relat_1(X5)|~v1_funct_1(X5)))&(X7=k1_funct_1(X5,esk1_3(X5,X6,X7))|~r2_hidden(X7,X6)|X6!=k2_relat_1(X5)|(~v1_relat_1(X5)|~v1_funct_1(X5))))&(~r2_hidden(X10,k1_relat_1(X5))|X9!=k1_funct_1(X5,X10)|r2_hidden(X9,X6)|X6!=k2_relat_1(X5)|(~v1_relat_1(X5)|~v1_funct_1(X5))))&((~r2_hidden(esk2_2(X5,X11),X11)|(~r2_hidden(X13,k1_relat_1(X5))|esk2_2(X5,X11)!=k1_funct_1(X5,X13))|X11=k2_relat_1(X5)|(~v1_relat_1(X5)|~v1_funct_1(X5)))&((r2_hidden(esk3_2(X5,X11),k1_relat_1(X5))|r2_hidden(esk2_2(X5,X11),X11)|X11=k2_relat_1(X5)|(~v1_relat_1(X5)|~v1_funct_1(X5)))&(esk2_2(X5,X11)=k1_funct_1(X5,esk3_2(X5,X11))|r2_hidden(esk2_2(X5,X11),X11)|X11=k2_relat_1(X5)|(~v1_relat_1(X5)|~v1_funct_1(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_funct_1])])])])])])).
fof(c_0_6, plain, ![X15, X16, X18]:((v1_relat_1(esk4_2(X15,X16))&v1_funct_1(esk4_2(X15,X16)))&(k1_relat_1(esk4_2(X15,X16))=X15&(~r2_hidden(X18,X15)|k1_funct_1(esk4_2(X15,X16),X18)=X16))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[s3_funct_1__e2_24__funct_1])])])])).
cnf(c_0_7, plain, (X1=k1_funct_1(X2,esk1_3(X2,X3,X1))|~r2_hidden(X1,X3)|X3!=k2_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8, plain, (r2_hidden(esk1_3(X1,X2,X3),k1_relat_1(X1))|~r2_hidden(X3,X2)|X2!=k2_relat_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (k1_funct_1(esk4_2(X2,X3),X1)=X3|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (k1_funct_1(X1,esk1_3(X1,k2_relat_1(X1),X2))=X2|~r2_hidden(X2,k2_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (v1_funct_1(esk4_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12, plain, (v1_relat_1(esk4_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13, plain, (r2_hidden(esk1_3(X1,k2_relat_1(X1),X2),k1_relat_1(X1))|~r2_hidden(X2,k2_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_8])).
cnf(c_0_14, plain, (k1_relat_1(esk4_2(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_15, negated_conjecture, ~(![X1]:(~(X1=k1_xboole_0)=>![X2]:?[X3]:((v1_relat_1(X3)&v1_funct_1(X3))&(k1_relat_1(X3)=X1&k2_relat_1(X3)=k1_tarski(X2))))), inference(assume_negation,[status(cth)],[t15_funct_1])).
cnf(c_0_16, plain, (X1=X2|~r2_hidden(esk1_3(esk4_2(X3,X2),k2_relat_1(esk4_2(X3,X2)),X1),X3)|~r2_hidden(X1,k2_relat_1(esk4_2(X3,X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11]), c_0_12])])).
cnf(c_0_17, plain, (r2_hidden(esk1_3(esk4_2(X1,X2),k2_relat_1(esk4_2(X1,X2)),X3),X1)|~r2_hidden(X3,k2_relat_1(esk4_2(X1,X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_11]), c_0_12])])).
fof(c_0_18, plain, ![X26, X27]:((r2_hidden(esk6_2(X26,X27),X26)|X26=k1_xboole_0|X26=k1_tarski(X27))&(esk6_2(X26,X27)!=X27|X26=k1_xboole_0|X26=k1_tarski(X27))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t41_zfmisc_1])])])])])).
fof(c_0_19, negated_conjecture, ![X31]:(esk7_0!=k1_xboole_0&(~v1_relat_1(X31)|~v1_funct_1(X31)|(k1_relat_1(X31)!=esk7_0|k2_relat_1(X31)!=k1_tarski(esk8_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_15])])])])])).
cnf(c_0_20, plain, (X1=X2|~r2_hidden(X1,k2_relat_1(esk4_2(X3,X2)))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_21, plain, (r2_hidden(esk6_2(X1,X2),X1)|X1=k1_xboole_0|X1=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, negated_conjecture, (~v1_relat_1(X1)|~v1_funct_1(X1)|k1_relat_1(X1)!=esk7_0|k2_relat_1(X1)!=k1_tarski(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|esk6_2(X1,X2)!=X2), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (esk6_2(k2_relat_1(esk4_2(X1,X2)),X3)=X2|k2_relat_1(esk4_2(X1,X2))=k1_tarski(X3)|k2_relat_1(esk4_2(X1,X2))=k1_xboole_0), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_25, plain, (r2_hidden(X3,X4)|~r2_hidden(X1,k1_relat_1(X2))|X3!=k1_funct_1(X2,X1)|X4!=k2_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_26, negated_conjecture, (k2_relat_1(esk4_2(esk7_0,X1))!=k1_tarski(esk8_0)), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_14]), c_0_11]), c_0_12])])])).
cnf(c_0_27, plain, (k2_relat_1(esk4_2(X1,X2))=k1_tarski(X2)|k2_relat_1(esk4_2(X1,X2))=k1_xboole_0), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24])])).
fof(c_0_28, plain, ![X19, X20, X21, X22, X23, X24]:(((~r2_hidden(X21,X20)|X21=X19|X20!=k1_tarski(X19))&(X22!=X19|r2_hidden(X22,X20)|X20!=k1_tarski(X19)))&((~r2_hidden(esk5_2(X23,X24),X24)|esk5_2(X23,X24)!=X23|X24=k1_tarski(X23))&(r2_hidden(esk5_2(X23,X24),X24)|esk5_2(X23,X24)=X23|X24=k1_tarski(X23)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_29, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relat_1(X1))|~r2_hidden(X2,k1_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_25])])).
cnf(c_0_30, negated_conjecture, (k2_relat_1(esk4_2(esk7_0,X1))=k1_xboole_0|k1_tarski(X1)!=k1_tarski(esk8_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_31, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_32, plain, (r2_hidden(X1,k2_relat_1(esk4_2(X2,X1)))|~r2_hidden(X3,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_9]), c_0_14]), c_0_11]), c_0_12])])).
cnf(c_0_33, negated_conjecture, (X1=X2|k1_tarski(X2)!=k1_tarski(esk8_0)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_20, c_0_30])).
cnf(c_0_34, plain, (r2_hidden(esk5_2(X1,X2),X2)|esk5_2(X1,X2)=X1|X2=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, plain, (r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_31])])).
cnf(c_0_36, plain, (X1=k1_tarski(X2)|X1=k1_xboole_0|r2_hidden(X3,k2_relat_1(esk4_2(X1,X3)))), inference(spm,[status(thm)],[c_0_32, c_0_21])).
cnf(c_0_37, negated_conjecture, (esk5_2(X1,k1_xboole_0)=X1|esk5_2(X1,k1_xboole_0)=X2|k1_tarski(X1)=k1_xboole_0|k1_tarski(X2)!=k1_tarski(esk8_0)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_38, negated_conjecture, (k1_tarski(X1)!=k1_tarski(esk8_0)|k1_tarski(esk8_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_26, c_0_30])).
cnf(c_0_39, plain, (X1=k1_xboole_0|r2_hidden(X2,k2_relat_1(esk4_2(X1,X2)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_32])).
cnf(c_0_40, negated_conjecture, (esk7_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_41, negated_conjecture, (esk5_2(X1,k1_xboole_0)=esk8_0|esk5_2(X1,k1_xboole_0)=X1|k1_tarski(X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_42, negated_conjecture, (k1_tarski(esk8_0)!=k1_xboole_0), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_43, negated_conjecture, (r2_hidden(X1,k1_xboole_0)|k1_tarski(X1)!=k1_tarski(esk8_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_30]), c_0_40])).
cnf(c_0_44, plain, (X2=k1_tarski(X1)|~r2_hidden(esk5_2(X1,X2),X2)|esk5_2(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_45, negated_conjecture, (esk5_2(esk8_0,k1_xboole_0)=esk8_0), inference(sr,[status(thm)],[inference(er,[status(thm)],[inference(ef,[status(thm)],[c_0_41])]), c_0_42])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk8_0,k1_xboole_0)), inference(er,[status(thm)],[c_0_43])).
cnf(c_0_47, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])]), c_0_42]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 48
# Proof object clause steps            : 37
# Proof object formula steps           : 11
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 28
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 18
# Processed clauses                    : 415
# ...of these trivial                  : 3
# ...subsumed                          : 257
# ...remaining for further processing  : 155
# Other redundant clauses eliminated   : 26
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 25
# Backward-rewritten                   : 51
# Generated clauses                    : 1188
# ...of the previous two non-trivial   : 1001
# Contextual simplify-reflections      : 2
# Paramodulations                      : 1147
# Factorizations                       : 8
# Equation resolutions                 : 35
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 56
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 45
# Current number of unprocessed clauses: 480
# ...number of literals in the above   : 1495
# Current number of archived formulas  : 0
# Current number of archived clauses   : 94
# Clause-clause subsumption calls (NU) : 3111
# Rec. Clause-clause subsumption calls : 2135
# Non-unit clause-clause subsumptions  : 269
# Unit Clause-clause subsumption calls : 134
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 144
# BW rewrite match successes           : 44
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 23281

# -------------------------------------------------
# User time                : 0.019 s
# System time              : 0.003 s
# Total time               : 0.022 s
# Maximum resident set size: 28712 pages
