# No SInE strategy applied
# Auto-Mode selected heuristic G_E___200_B02_F1_SE_CS_SP_PI_S2k
# and selection function SelectCQArNTEqFirst.
#
# Preprocessing time       : 0.010 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_boole, axiom, ![X1, X2]:~((v1_xboole_0(X1)&(~(X1=X2)&v1_xboole_0(X2)))), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', t8_boole)).
fof(cc4_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', cc4_relset_1)).
fof(dt_k6_partfun1, axiom, ![X1]:(v1_partfun1(k6_partfun1(X1),X1)&m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', dt_k6_partfun1)).
fof(t76_funct_2, conjecture, ![X1, X2]:((v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>((r2_relset_1(X1,X1,k1_partfun1(X1,X1,X1,X1,X3,X2),X2)&v2_funct_1(X2))=>r2_relset_1(X1,X1,X3,k6_partfun1(X1))))), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', t76_funct_2)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', fc1_xboole_0)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', d10_xboole_0)).
fof(t67_funct_2, axiom, ![X1, X2]:((v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>k1_relset_1(X1,X1,X2)=X1), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', t67_funct_2)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', t113_zfmisc_1)).
fof(fc4_funct_1, axiom, ![X1]:(v1_relat_1(k6_relat_1(X1))&v2_funct_1(k6_relat_1(X1))), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', fc4_funct_1)).
fof(redefinition_k6_partfun1, axiom, ![X1]:k6_partfun1(X1)=k6_relat_1(X1), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', redefinition_k6_partfun1)).
fof(t4_funct_2, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r1_tarski(k2_relat_1(X2),X1)=>(v1_funct_1(X2)&(v1_funct_2(X2,k1_relat_1(X2),X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X2),X1))))))), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', t4_funct_2)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', redefinition_k1_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', cc1_relset_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', cc2_relset_1)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', redefinition_k2_relset_1)).
fof(d19_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(v5_relat_1(X2,X1)<=>r1_tarski(k2_relat_1(X2),X1))), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', d19_relat_1)).
fof(redefinition_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:(((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&(v1_funct_1(X6)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))))=>k1_partfun1(X1,X2,X3,X4,X5,X6)=k5_relat_1(X5,X6)), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', redefinition_k1_partfun1)).
fof(t31_funct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>((v2_funct_1(X3)&k2_relset_1(X1,X2,X3)=X2)=>(v1_funct_1(k2_funct_1(X3))&(v1_funct_2(k2_funct_1(X3),X2,X1)&m1_subset_1(k2_funct_1(X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1))))))), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', t31_funct_2)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', t3_subset)).
fof(redefinition_r2_relset_1, axiom, ![X1, X2, X3, X4]:((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>(r2_relset_1(X1,X2,X3,X4)<=>X3=X4)), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', redefinition_r2_relset_1)).
fof(dt_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:(((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&(v1_funct_1(X6)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))))=>(v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))&m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4))))), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', dt_k1_partfun1)).
fof(t71_relat_1, axiom, ![X1]:(k1_relat_1(k6_relat_1(X1))=X1&k2_relat_1(k6_relat_1(X1))=X1), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', t71_relat_1)).
fof(t61_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)=>(k5_relat_1(X1,k2_funct_1(X1))=k6_relat_1(k1_relat_1(X1))&k5_relat_1(k2_funct_1(X1),X1)=k6_relat_1(k2_relat_1(X1))))), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', t61_funct_1)).
fof(redefinition_k4_relset_1, axiom, ![X1, X2, X3, X4, X5, X6]:((m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>k4_relset_1(X1,X2,X3,X4,X5,X6)=k5_relat_1(X5,X6)), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', redefinition_k4_relset_1)).
fof(t23_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(r2_relset_1(X1,X2,k4_relset_1(X1,X1,X1,X2,k6_partfun1(X1),X3),X3)&r2_relset_1(X1,X2,k4_relset_1(X1,X2,X2,X2,X3,k6_partfun1(X2)),X3))), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', t23_funct_2)).
fof(t27_funct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&~((v2_funct_1(X3)<=>![X4, X5]:((v1_funct_1(X5)&(v1_funct_2(X5,X4,X1)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X1)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,X4,X1)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X1)))))=>(r2_relset_1(X4,X2,k1_partfun1(X4,X1,X1,X2,X5,X3),k1_partfun1(X4,X1,X1,X2,X6,X3))=>r2_relset_1(X4,X1,X5,X6)))))))))), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', t27_funct_2)).
fof(symmetry_r2_relset_1, axiom, ![X1, X2, X3, X4]:((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>(r2_relset_1(X1,X2,X3,X4)=>r2_relset_1(X1,X2,X4,X3))), file('data/MPTP/proofs/4a8ff414330f470d8ee2e6d6b34c1a61.E_input', symmetry_r2_relset_1)).
fof(c_0_27, plain, ![X24, X25]:(~v1_xboole_0(X24)|(X24=X25|~v1_xboole_0(X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_boole])])])).
fof(c_0_28, plain, ![X10, X11, X12]:(~v1_xboole_0(X10)|(~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(X11,X10)))|v1_xboole_0(X12))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])).
fof(c_0_29, plain, ![X56]:(v1_partfun1(k6_partfun1(X56),X56)&m1_subset_1(k6_partfun1(X56),k1_zfmisc_1(k2_zfmisc_1(X56,X56)))), inference(variable_rename,[status(thm)],[dt_k6_partfun1])).
fof(c_0_30, negated_conjecture, ~(![X1, X2]:((v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>((r2_relset_1(X1,X1,k1_partfun1(X1,X1,X1,X1,X3,X2),X2)&v2_funct_1(X2))=>r2_relset_1(X1,X1,X3,k6_partfun1(X1)))))), inference(assume_negation,[status(cth)],[t76_funct_2])).
cnf(c_0_31, plain, (X1=X2|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_32, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_33, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34, plain, (m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_35, plain, ![X82, X83]:(((r1_tarski(X82,X83)|X82!=X83)&(r1_tarski(X83,X82)|X82!=X83))&(~r1_tarski(X82,X83)|~r1_tarski(X83,X82)|X82=X83)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_36, plain, ![X53, X54]:(~v1_funct_1(X54)|(~v1_funct_2(X54,X53,X53)|~m1_subset_1(X54,k1_zfmisc_1(k2_zfmisc_1(X53,X53))))|k1_relset_1(X53,X53,X54)=X53), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t67_funct_2])])).
fof(c_0_37, negated_conjecture, ((v1_funct_1(esk2_0)&(v1_funct_2(esk2_0,esk1_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,esk1_0,esk1_0)&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))))&((r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0),esk2_0)&v2_funct_1(esk2_0))&~r2_relset_1(esk1_0,esk1_0,esk3_0,k6_partfun1(esk1_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])).
fof(c_0_38, plain, ![X39, X40]:((k2_zfmisc_1(X39,X40)!=k1_xboole_0|(X39=k1_xboole_0|X40=k1_xboole_0))&((X39!=k1_xboole_0|k2_zfmisc_1(X39,X40)=k1_xboole_0)&(X40!=k1_xboole_0|k2_zfmisc_1(X39,X40)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_39, plain, (k1_xboole_0=X1|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_40, plain, (v1_xboole_0(k6_partfun1(X1))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
fof(c_0_41, plain, ![X55]:(v1_relat_1(k6_relat_1(X55))&v2_funct_1(k6_relat_1(X55))), inference(variable_rename,[status(thm)],[fc4_funct_1])).
fof(c_0_42, plain, ![X84]:k6_partfun1(X84)=k6_relat_1(X84), inference(variable_rename,[status(thm)],[redefinition_k6_partfun1])).
fof(c_0_43, plain, ![X57, X58]:((v1_funct_1(X58)|~r1_tarski(k2_relat_1(X58),X57)|(~v1_relat_1(X58)|~v1_funct_1(X58)))&((v1_funct_2(X58,k1_relat_1(X58),X57)|~r1_tarski(k2_relat_1(X58),X57)|(~v1_relat_1(X58)|~v1_funct_1(X58)))&(m1_subset_1(X58,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X58),X57)))|~r1_tarski(k2_relat_1(X58),X57)|(~v1_relat_1(X58)|~v1_funct_1(X58))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_funct_2])])])).
cnf(c_0_44, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_45, plain, ![X44, X45, X46]:(~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(X44,X45)))|k1_relset_1(X44,X45,X46)=k1_relat_1(X46)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_46, plain, (k1_relset_1(X2,X2,X1)=X2|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_47, negated_conjecture, (v1_funct_2(esk2_0,esk1_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_48, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_49, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_50, plain, ![X36, X37, X38]:(~m1_subset_1(X38,k1_zfmisc_1(k2_zfmisc_1(X36,X37)))|v1_relat_1(X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_51, plain, ![X18, X19, X20]:((v4_relat_1(X20,X18)|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(X18,X19))))&(v5_relat_1(X20,X19)|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(X18,X19))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_52, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_53, plain, (k6_partfun1(X1)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_54, plain, (v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_55, plain, (k6_partfun1(X1)=k6_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_56, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),X2)))|~r1_tarski(k2_relat_1(X1),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_57, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_58, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_59, negated_conjecture, (k1_relset_1(esk1_0,esk1_0,esk2_0)=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48]), c_0_49])])).
cnf(c_0_60, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_61, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_62, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_52])).
cnf(c_0_63, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~v1_xboole_0(k6_partfun1(X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_34, c_0_31])).
cnf(c_0_64, plain, (k6_partfun1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_53, c_0_32])).
cnf(c_0_65, plain, (v1_relat_1(k6_partfun1(X1))), inference(rw,[status(thm)],[c_0_54, c_0_55])).
fof(c_0_66, plain, ![X21, X22, X23]:(~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(X21,X22)))|k2_relset_1(X21,X22,X23)=k2_relat_1(X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
cnf(c_0_67, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_68, negated_conjecture, (k1_relat_1(esk2_0)=esk1_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_48]), c_0_59])).
cnf(c_0_69, negated_conjecture, (v1_relat_1(esk2_0)), inference(spm,[status(thm)],[c_0_60, c_0_48])).
cnf(c_0_70, plain, (v1_funct_2(X1,k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_71, plain, ![X76, X77]:((~v5_relat_1(X77,X76)|r1_tarski(k2_relat_1(X77),X76)|~v1_relat_1(X77))&(~r1_tarski(k2_relat_1(X77),X76)|v5_relat_1(X77,X76)|~v1_relat_1(X77))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d19_relat_1])])])).
cnf(c_0_72, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_73, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_62]), c_0_32])])).
cnf(c_0_74, plain, (v1_relat_1(X1)|~v1_xboole_0(k6_partfun1(X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_65, c_0_31])).
fof(c_0_75, plain, ![X26, X27, X28, X29, X30, X31]:(~v1_funct_1(X30)|~m1_subset_1(X30,k1_zfmisc_1(k2_zfmisc_1(X26,X27)))|(~v1_funct_1(X31)|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(X28,X29))))|k1_partfun1(X26,X27,X28,X29,X30,X31)=k5_relat_1(X30,X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_partfun1])])).
fof(c_0_76, plain, ![X13, X14, X15]:((v1_funct_1(k2_funct_1(X15))|(~v2_funct_1(X15)|k2_relset_1(X13,X14,X15)!=X14)|(~v1_funct_1(X15)|(~v1_funct_2(X15,X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))))&((v1_funct_2(k2_funct_1(X15),X14,X13)|(~v2_funct_1(X15)|k2_relset_1(X13,X14,X15)!=X14)|(~v1_funct_1(X15)|(~v1_funct_2(X15,X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))))&(m1_subset_1(k2_funct_1(X15),k1_zfmisc_1(k2_zfmisc_1(X14,X13)))|(~v2_funct_1(X15)|k2_relset_1(X13,X14,X15)!=X14)|(~v1_funct_1(X15)|(~v1_funct_2(X15,X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_funct_2])])])).
cnf(c_0_77, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_78, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k2_relat_1(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_69]), c_0_49])])).
cnf(c_0_79, plain, (v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_70, c_0_57])).
fof(c_0_80, plain, ![X16, X17]:((~m1_subset_1(X16,k1_zfmisc_1(X17))|r1_tarski(X16,X17))&(~r1_tarski(X16,X17)|m1_subset_1(X16,k1_zfmisc_1(X17)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_81, plain, (r1_tarski(k2_relat_1(X1),X2)|~v5_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_82, plain, (v5_relat_1(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_83, plain, (v1_relat_1(X1)|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_64]), c_0_32])])).
cnf(c_0_84, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_85, plain, ![X78, X79, X80, X81]:((~r2_relset_1(X78,X79,X80,X81)|X80=X81|(~m1_subset_1(X80,k1_zfmisc_1(k2_zfmisc_1(X78,X79)))|~m1_subset_1(X81,k1_zfmisc_1(k2_zfmisc_1(X78,X79)))))&(X80!=X81|r2_relset_1(X78,X79,X80,X81)|(~m1_subset_1(X80,k1_zfmisc_1(k2_zfmisc_1(X78,X79)))|~m1_subset_1(X81,k1_zfmisc_1(k2_zfmisc_1(X78,X79)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r2_relset_1])])])).
cnf(c_0_86, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k5_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_75])).
fof(c_0_87, plain, ![X47, X48, X49, X50, X51, X52]:((v1_funct_1(k1_partfun1(X47,X48,X49,X50,X51,X52))|(~v1_funct_1(X51)|~m1_subset_1(X51,k1_zfmisc_1(k2_zfmisc_1(X47,X48)))|(~v1_funct_1(X52)|~m1_subset_1(X52,k1_zfmisc_1(k2_zfmisc_1(X49,X50))))))&(m1_subset_1(k1_partfun1(X47,X48,X49,X50,X51,X52),k1_zfmisc_1(k2_zfmisc_1(X47,X50)))|(~v1_funct_1(X51)|~m1_subset_1(X51,k1_zfmisc_1(k2_zfmisc_1(X47,X48)))|(~v1_funct_1(X52)|~m1_subset_1(X52,k1_zfmisc_1(k2_zfmisc_1(X49,X50))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_partfun1])])])).
cnf(c_0_88, plain, (m1_subset_1(k2_funct_1(X1),k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v2_funct_1(X1)|k2_relset_1(X3,X2,X1)!=X2|~v1_funct_1(X1)|~v1_funct_2(X1,X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_89, negated_conjecture, (k2_relset_1(esk1_0,k2_relat_1(esk2_0),esk2_0)=k2_relat_1(esk2_0)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_90, negated_conjecture, (v2_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_91, negated_conjecture, (v1_funct_2(esk2_0,esk1_0,k2_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_68]), c_0_69]), c_0_49])])).
cnf(c_0_92, plain, (v1_funct_1(k2_funct_1(X1))|~v2_funct_1(X1)|k2_relset_1(X2,X3,X1)!=X3|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_93, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_94, plain, (r1_tarski(k2_relat_1(X1),X2)|~v1_xboole_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_83])).
cnf(c_0_95, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_84])).
cnf(c_0_96, plain, (X3=X4|~r2_relset_1(X1,X2,X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_97, negated_conjecture, (r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_98, negated_conjecture, (k1_partfun1(X1,X2,esk1_0,esk1_0,X3,esk2_0)=k5_relat_1(X3,esk2_0)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_48]), c_0_49])])).
cnf(c_0_99, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_100, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_101, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_102, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_103, negated_conjecture, (m1_subset_1(k2_funct_1(esk2_0),k1_zfmisc_1(k2_zfmisc_1(k2_relat_1(esk2_0),esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_90]), c_0_78]), c_0_91]), c_0_49])])).
cnf(c_0_104, negated_conjecture, (v1_funct_1(k2_funct_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_89]), c_0_90]), c_0_78]), c_0_91]), c_0_49])])).
fof(c_0_105, plain, ![X75]:(k1_relat_1(k6_relat_1(X75))=X75&k2_relat_1(k6_relat_1(X75))=X75), inference(variable_rename,[status(thm)],[t71_relat_1])).
cnf(c_0_106, plain, (m1_subset_1(k2_relat_1(X1),k1_zfmisc_1(X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_107, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_60, c_0_95])).
cnf(c_0_108, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)=esk2_0|~m1_subset_1(k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_48])])).
cnf(c_0_109, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)=k5_relat_1(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_100])])).
cnf(c_0_110, negated_conjecture, (m1_subset_1(k1_partfun1(X1,X2,esk1_0,esk1_0,X3,esk2_0),k1_zfmisc_1(k2_zfmisc_1(X1,esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_48]), c_0_49])])).
cnf(c_0_111, negated_conjecture, (v1_funct_1(k1_partfun1(X1,X2,k2_relat_1(esk2_0),esk1_0,X3,k2_funct_1(esk2_0)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_104])])).
cnf(c_0_112, negated_conjecture, (k1_partfun1(X1,X2,k2_relat_1(esk2_0),esk1_0,X3,k2_funct_1(esk2_0))=k5_relat_1(X3,k2_funct_1(esk2_0))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_103]), c_0_104])])).
fof(c_0_113, plain, ![X65]:((k5_relat_1(X65,k2_funct_1(X65))=k6_relat_1(k1_relat_1(X65))|~v2_funct_1(X65)|(~v1_relat_1(X65)|~v1_funct_1(X65)))&(k5_relat_1(k2_funct_1(X65),X65)=k6_relat_1(k2_relat_1(X65))|~v2_funct_1(X65)|(~v1_relat_1(X65)|~v1_funct_1(X65)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t61_funct_1])])])).
fof(c_0_114, plain, ![X59, X60, X61, X62, X63, X64]:(~m1_subset_1(X63,k1_zfmisc_1(k2_zfmisc_1(X59,X60)))|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(X61,X62)))|k4_relset_1(X59,X60,X61,X62,X63,X64)=k5_relat_1(X63,X64)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_relset_1])])).
cnf(c_0_115, plain, (k2_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_116, plain, (k1_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_117, plain, (v5_relat_1(k2_relat_1(X1),X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_72, c_0_106])).
cnf(c_0_118, plain, (v1_relat_1(k2_relat_1(X1))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_107, c_0_106])).
cnf(c_0_119, negated_conjecture, (v1_funct_2(esk3_0,esk1_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_120, negated_conjecture, (k5_relat_1(esk3_0,esk2_0)=esk2_0|~m1_subset_1(k5_relat_1(esk3_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_108, c_0_109]), c_0_109])).
cnf(c_0_121, negated_conjecture, (m1_subset_1(k5_relat_1(esk3_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_99]), c_0_109]), c_0_100])])).
cnf(c_0_122, negated_conjecture, (v1_funct_1(k1_partfun1(esk1_0,esk1_0,k2_relat_1(esk2_0),esk1_0,esk2_0,k2_funct_1(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_48]), c_0_49])])).
cnf(c_0_123, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,k2_relat_1(esk2_0),esk1_0,esk2_0,k2_funct_1(esk2_0))=k5_relat_1(esk2_0,k2_funct_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_48]), c_0_49])])).
cnf(c_0_124, plain, (k5_relat_1(X1,k2_funct_1(X1))=k6_relat_1(k1_relat_1(X1))|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_113])).
fof(c_0_125, plain, ![X41, X42, X43]:((r2_relset_1(X41,X42,k4_relset_1(X41,X41,X41,X42,k6_partfun1(X41),X43),X43)|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(X41,X42))))&(r2_relset_1(X41,X42,k4_relset_1(X41,X42,X42,X42,X43,k6_partfun1(X42)),X43)|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(X41,X42))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_funct_2])])])).
cnf(c_0_126, plain, (k4_relset_1(X2,X3,X5,X6,X1,X4)=k5_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_114])).
cnf(c_0_127, plain, (k2_relat_1(k6_partfun1(X1))=X1), inference(rw,[status(thm)],[c_0_115, c_0_55])).
cnf(c_0_128, plain, (k1_relat_1(k6_partfun1(X1))=X1), inference(rw,[status(thm)],[c_0_116, c_0_55])).
cnf(c_0_129, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_130, plain, (r1_tarski(k2_relat_1(k2_relat_1(X1)),X2)|~v1_xboole_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_117]), c_0_118])).
cnf(c_0_131, negated_conjecture, (k1_relset_1(esk1_0,esk1_0,esk3_0)=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_119]), c_0_99]), c_0_100])])).
fof(c_0_132, plain, ![X66, X67, X68, X69, X70, X71]:((~v2_funct_1(X68)|(~v1_funct_1(X70)|(~v1_funct_2(X70,X69,X66)|~m1_subset_1(X70,k1_zfmisc_1(k2_zfmisc_1(X69,X66))))|(~v1_funct_1(X71)|(~v1_funct_2(X71,X69,X66)|~m1_subset_1(X71,k1_zfmisc_1(k2_zfmisc_1(X69,X66))))|(~r2_relset_1(X69,X67,k1_partfun1(X69,X66,X66,X67,X70,X68),k1_partfun1(X69,X66,X66,X67,X71,X68))|r2_relset_1(X69,X66,X70,X71))))|X67=k1_xboole_0|X66=k1_xboole_0|(~v1_funct_1(X68)|(~v1_funct_2(X68,X66,X67)|~m1_subset_1(X68,k1_zfmisc_1(k2_zfmisc_1(X66,X67))))))&(((v1_funct_1(esk5_3(X66,X67,X68))|v2_funct_1(X68)|X67=k1_xboole_0|X66=k1_xboole_0|(~v1_funct_1(X68)|(~v1_funct_2(X68,X66,X67)|~m1_subset_1(X68,k1_zfmisc_1(k2_zfmisc_1(X66,X67))))))&((v1_funct_2(esk5_3(X66,X67,X68),esk4_3(X66,X67,X68),X66)|v2_funct_1(X68)|X67=k1_xboole_0|X66=k1_xboole_0|(~v1_funct_1(X68)|(~v1_funct_2(X68,X66,X67)|~m1_subset_1(X68,k1_zfmisc_1(k2_zfmisc_1(X66,X67))))))&(m1_subset_1(esk5_3(X66,X67,X68),k1_zfmisc_1(k2_zfmisc_1(esk4_3(X66,X67,X68),X66)))|v2_funct_1(X68)|X67=k1_xboole_0|X66=k1_xboole_0|(~v1_funct_1(X68)|(~v1_funct_2(X68,X66,X67)|~m1_subset_1(X68,k1_zfmisc_1(k2_zfmisc_1(X66,X67))))))))&(((v1_funct_1(esk6_3(X66,X67,X68))|v2_funct_1(X68)|X67=k1_xboole_0|X66=k1_xboole_0|(~v1_funct_1(X68)|(~v1_funct_2(X68,X66,X67)|~m1_subset_1(X68,k1_zfmisc_1(k2_zfmisc_1(X66,X67))))))&((v1_funct_2(esk6_3(X66,X67,X68),esk4_3(X66,X67,X68),X66)|v2_funct_1(X68)|X67=k1_xboole_0|X66=k1_xboole_0|(~v1_funct_1(X68)|(~v1_funct_2(X68,X66,X67)|~m1_subset_1(X68,k1_zfmisc_1(k2_zfmisc_1(X66,X67))))))&(m1_subset_1(esk6_3(X66,X67,X68),k1_zfmisc_1(k2_zfmisc_1(esk4_3(X66,X67,X68),X66)))|v2_funct_1(X68)|X67=k1_xboole_0|X66=k1_xboole_0|(~v1_funct_1(X68)|(~v1_funct_2(X68,X66,X67)|~m1_subset_1(X68,k1_zfmisc_1(k2_zfmisc_1(X66,X67))))))))&((r2_relset_1(esk4_3(X66,X67,X68),X67,k1_partfun1(esk4_3(X66,X67,X68),X66,X66,X67,esk5_3(X66,X67,X68),X68),k1_partfun1(esk4_3(X66,X67,X68),X66,X66,X67,esk6_3(X66,X67,X68),X68))|v2_funct_1(X68)|X67=k1_xboole_0|X66=k1_xboole_0|(~v1_funct_1(X68)|(~v1_funct_2(X68,X66,X67)|~m1_subset_1(X68,k1_zfmisc_1(k2_zfmisc_1(X66,X67))))))&(~r2_relset_1(esk4_3(X66,X67,X68),X66,esk5_3(X66,X67,X68),esk6_3(X66,X67,X68))|v2_funct_1(X68)|X67=k1_xboole_0|X66=k1_xboole_0|(~v1_funct_1(X68)|(~v1_funct_2(X68,X66,X67)|~m1_subset_1(X68,k1_zfmisc_1(k2_zfmisc_1(X66,X67)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t27_funct_2])])])])])])).
cnf(c_0_133, negated_conjecture, (k5_relat_1(esk3_0,esk2_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_120, c_0_121])])).
cnf(c_0_134, negated_conjecture, (v1_funct_1(k5_relat_1(esk2_0,k2_funct_1(esk2_0)))), inference(rw,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_135, plain, (k5_relat_1(X1,k2_funct_1(X1))=k6_partfun1(k1_relat_1(X1))|~v1_funct_1(X1)|~v2_funct_1(X1)|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_124, c_0_55])).
cnf(c_0_136, plain, (r2_relset_1(X1,X2,k4_relset_1(X1,X1,X1,X2,k6_partfun1(X1),X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_125])).
cnf(c_0_137, negated_conjecture, (k4_relset_1(X1,X2,esk1_0,esk1_0,X3,esk2_0)=k5_relat_1(X3,esk2_0)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_126, c_0_48])).
cnf(c_0_138, plain, (v1_funct_2(k6_partfun1(X1),X1,X2)|~r1_tarski(X1,X2)|~v1_funct_1(k6_partfun1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_127]), c_0_128]), c_0_65])])).
cnf(c_0_139, plain, (X1=k2_relat_1(k2_relat_1(X2))|~r1_tarski(X1,k2_relat_1(k2_relat_1(X2)))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_129, c_0_130])).
cnf(c_0_140, negated_conjecture, (k1_relat_1(esk3_0)=esk1_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_99]), c_0_131])).
cnf(c_0_141, negated_conjecture, (v1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_60, c_0_99])).
cnf(c_0_142, plain, (r2_relset_1(X3,X4,X2,X5)|X6=k1_xboole_0|X4=k1_xboole_0|~v2_funct_1(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X5)|~v1_funct_2(X5,X3,X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~r2_relset_1(X3,X6,k1_partfun1(X3,X4,X4,X6,X2,X1),k1_partfun1(X3,X4,X4,X6,X5,X1))|~v1_funct_1(X1)|~v1_funct_2(X1,X4,X6)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X6)))), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_143, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)=esk2_0), inference(rw,[status(thm)],[c_0_109, c_0_133])).
cnf(c_0_144, negated_conjecture, (k1_partfun1(X1,X1,esk1_0,esk1_0,k6_partfun1(X1),esk2_0)=k5_relat_1(k6_partfun1(X1),esk2_0)|~v1_funct_1(k6_partfun1(X1))), inference(spm,[status(thm)],[c_0_98, c_0_34])).
cnf(c_0_145, negated_conjecture, (v1_funct_1(k6_partfun1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_68]), c_0_69]), c_0_90]), c_0_49])])).
cnf(c_0_146, negated_conjecture, (r2_relset_1(esk1_0,esk1_0,k4_relset_1(esk1_0,esk1_0,esk1_0,esk1_0,k6_partfun1(esk1_0),esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_136, c_0_48])).
cnf(c_0_147, negated_conjecture, (k4_relset_1(X1,X1,esk1_0,esk1_0,k6_partfun1(X1),esk2_0)=k5_relat_1(k6_partfun1(X1),esk2_0)), inference(spm,[status(thm)],[c_0_137, c_0_34])).
cnf(c_0_148, plain, (v1_funct_2(k6_partfun1(X1),X1,X1)|~v1_funct_1(k6_partfun1(X1))), inference(spm,[status(thm)],[c_0_138, c_0_57])).
cnf(c_0_149, plain, (X1=k2_relat_1(X2)|~r1_tarski(X1,k2_relat_1(X2))|~v1_xboole_0(k6_partfun1(X2))), inference(spm,[status(thm)],[c_0_139, c_0_127])).
cnf(c_0_150, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_151, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_140]), c_0_141]), c_0_100])])).
fof(c_0_152, plain, ![X32, X33, X34, X35]:(~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33)))|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(X32,X33)))|(~r2_relset_1(X32,X33,X34,X35)|r2_relset_1(X32,X33,X35,X34))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r2_relset_1])])).
cnf(c_0_153, negated_conjecture, (esk1_0=k1_xboole_0|r2_relset_1(esk1_0,esk1_0,X1,esk3_0)|~r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))|~v1_funct_2(X1,esk1_0,esk1_0)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_143]), c_0_90]), c_0_99]), c_0_48]), c_0_119]), c_0_47]), c_0_100]), c_0_49])])).
cnf(c_0_154, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,k6_partfun1(esk1_0),esk2_0)=k5_relat_1(k6_partfun1(esk1_0),esk2_0)), inference(spm,[status(thm)],[c_0_144, c_0_145])).
cnf(c_0_155, negated_conjecture, (r2_relset_1(esk1_0,esk1_0,k5_relat_1(k6_partfun1(esk1_0),esk2_0),esk2_0)), inference(rw,[status(thm)],[c_0_146, c_0_147])).
cnf(c_0_156, negated_conjecture, (v1_funct_2(k6_partfun1(esk1_0),esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_148, c_0_145])).
cnf(c_0_157, plain, (X1=X2|~r1_tarski(X1,X2)|~v1_xboole_0(k6_partfun1(k6_partfun1(X2)))), inference(spm,[status(thm)],[c_0_149, c_0_127])).
cnf(c_0_158, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0)))), inference(spm,[status(thm)],[c_0_150, c_0_151])).
cnf(c_0_159, plain, (r2_relset_1(X2,X3,X4,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r2_relset_1(X2,X3,X1,X4)), inference(split_conjunct,[status(thm)],[c_0_152])).
cnf(c_0_160, negated_conjecture, (esk1_0=k1_xboole_0|r2_relset_1(esk1_0,esk1_0,k6_partfun1(esk1_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_154]), c_0_155]), c_0_34]), c_0_156]), c_0_145])])).
cnf(c_0_161, negated_conjecture, (~r2_relset_1(esk1_0,esk1_0,esk3_0,k6_partfun1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_162, negated_conjecture, (r1_tarski(esk2_0,k2_zfmisc_1(esk1_0,esk1_0))), inference(spm,[status(thm)],[c_0_150, c_0_48])).
cnf(c_0_163, negated_conjecture, (k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0))=esk3_0|~v1_xboole_0(k6_partfun1(k6_partfun1(k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0)))))), inference(spm,[status(thm)],[c_0_157, c_0_158])).
cnf(c_0_164, negated_conjecture, (esk1_0=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_160]), c_0_99]), c_0_34])]), c_0_161])).
cnf(c_0_165, plain, (r2_relset_1(X3,X4,X1,X2)|X1!=X2|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_166, negated_conjecture, (k2_zfmisc_1(esk1_0,esk1_0)=esk2_0|~v1_xboole_0(k6_partfun1(k6_partfun1(k2_zfmisc_1(esk1_0,esk1_0))))), inference(spm,[status(thm)],[c_0_157, c_0_162])).
cnf(c_0_167, negated_conjecture, (k1_xboole_0=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_163, c_0_164]), c_0_62]), c_0_164]), c_0_62]), c_0_64]), c_0_64]), c_0_32])])).
cnf(c_0_168, plain, (r2_relset_1(X1,X2,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(er,[status(thm)],[c_0_165])).
cnf(c_0_169, negated_conjecture, (esk3_0=esk2_0), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_166, c_0_164]), c_0_164]), c_0_62]), c_0_164]), c_0_164]), c_0_62]), c_0_64]), c_0_64]), c_0_32])]), c_0_167])).
cnf(c_0_170, plain, (r2_relset_1(X1,X1,k6_partfun1(X1),k6_partfun1(X1))), inference(spm,[status(thm)],[c_0_168, c_0_34])).
cnf(c_0_171, plain, (k6_partfun1(esk2_0)=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_167]), c_0_167]), c_0_169]), c_0_169])).
cnf(c_0_172, plain, (r2_relset_1(esk2_0,esk2_0,esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_170, c_0_171])).
cnf(c_0_173, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_161, c_0_164]), c_0_164]), c_0_164]), c_0_64]), c_0_167]), c_0_169]), c_0_167]), c_0_169]), c_0_169]), c_0_167]), c_0_169]), c_0_172])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 174
# Proof object clause steps            : 120
# Proof object formula steps           : 54
# Proof object conjectures             : 55
# Proof object clause conjectures      : 52
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 43
# Proof object initial formulas used   : 27
# Proof object generating inferences   : 60
# Proof object simplifying inferences  : 129
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 27
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 61
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 59
# Processed clauses                    : 24656
# ...of these trivial                  : 225
# ...subsumed                          : 16653
# ...remaining for further processing  : 7777
# Other redundant clauses eliminated   : 191
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 200
# Backward-rewritten                   : 6542
# Generated clauses                    : 280727
# ...of the previous two non-trivial   : 277476
# Contextual simplify-reflections      : 232
# Paramodulations                      : 280536
# Factorizations                       : 0
# Equation resolutions                 : 191
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1030
#    Positive orientable unit clauses  : 64
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 966
# Current number of unprocessed clauses: 250136
# ...number of literals in the above   : 1026833
# Current number of archived formulas  : 0
# Current number of archived clauses   : 6743
# Clause-clause subsumption calls (NU) : 6227065
# Rec. Clause-clause subsumption calls : 3544720
# Non-unit clause-clause subsumptions  : 17085
# Unit Clause-clause subsumption calls : 220027
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 86692
# BW rewrite match successes           : 66
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7025372

# -------------------------------------------------
# User time                : 4.465 s
# System time              : 0.116 s
# Total time               : 4.581 s
# Maximum resident set size: 183548 pages
