# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l48_waybel_9, axiom, ![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((X3=X4&(![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X5),X1,X1))&(v10_waybel_0(X2,X1)&r3_waybel_9(X1,X2,X3))))=>r2_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X4)))))), file('data/MPTP/proofs/4aca218bf47b48a0a237b733689dc87b.E_input', l48_waybel_9)).
fof(t35_waybel_9, conjecture, ![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((~(v2_struct_0(X3))&(v4_orders_2(X3)&(v7_waybel_0(X3)&l1_waybel_0(X3,X1))))=>((![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X4),X1,X1))&(v10_waybel_0(X3,X1)&r3_waybel_9(X1,X3,X2)))=>X2=k1_waybel_2(X1,X3))))), file('data/MPTP/proofs/4aca218bf47b48a0a237b733689dc87b.E_input', t35_waybel_9)).
fof(l49_waybel_9, axiom, ![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((X3=X4&(![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X5),X1,X1))&r3_waybel_9(X1,X2,X3)))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r2_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X5)=>r3_orders_2(X1,X4,X5)))))))), file('data/MPTP/proofs/4aca218bf47b48a0a237b733689dc87b.E_input', l49_waybel_9)).
fof(dt_l1_waybel_9, axiom, ![X1]:(l1_waybel_9(X1)=>(l1_pre_topc(X1)&l1_orders_2(X1))), file('data/MPTP/proofs/4aca218bf47b48a0a237b733689dc87b.E_input', dt_l1_waybel_9)).
fof(t15_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4))))))), file('data/MPTP/proofs/4aca218bf47b48a0a237b733689dc87b.E_input', t15_yellow_0)).
fof(cc1_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v1_lattice3(X1)=>~(v2_struct_0(X1)))), file('data/MPTP/proofs/4aca218bf47b48a0a237b733689dc87b.E_input', cc1_lattice3)).
fof(d1_waybel_2, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(l1_waybel_0(X2,X1)=>k1_waybel_2(X1,X2)=k4_yellow_2(X1,u1_waybel_0(X1,X2)))), file('data/MPTP/proofs/4aca218bf47b48a0a237b733689dc87b.E_input', d1_waybel_2)).
fof(dt_u1_waybel_0, axiom, ![X1, X2]:((l1_struct_0(X1)&l1_waybel_0(X2,X1))=>(v1_funct_1(u1_waybel_0(X1,X2))&(v1_funct_2(u1_waybel_0(X1,X2),u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(u1_waybel_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))), file('data/MPTP/proofs/4aca218bf47b48a0a237b733689dc87b.E_input', dt_u1_waybel_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/4aca218bf47b48a0a237b733689dc87b.E_input', dt_l1_orders_2)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/4aca218bf47b48a0a237b733689dc87b.E_input', redefinition_r3_orders_2)).
fof(d5_yellow_2, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(v1_relat_1(X2)=>k4_yellow_2(X1,X2)=k1_yellow_0(X1,k2_relat_1(X2)))), file('data/MPTP/proofs/4aca218bf47b48a0a237b733689dc87b.E_input', d5_yellow_2)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/4aca218bf47b48a0a237b733689dc87b.E_input', cc1_relset_1)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/4aca218bf47b48a0a237b733689dc87b.E_input', redefinition_k2_relset_1)).
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/4aca218bf47b48a0a237b733689dc87b.E_input', d9_yellow_0)).
fof(c_0_14, plain, ![X6, X7, X8, X9]:((m1_subset_1(esk1_4(X6,X7,X8,X9),u1_struct_0(X6))|(~v10_waybel_0(X7,X6)|~r3_waybel_9(X6,X7,X8))|X8!=X9|r2_lattice3(X6,k2_relset_1(u1_struct_0(X7),u1_struct_0(X6),u1_waybel_0(X6,X7)),X9)|~m1_subset_1(X9,u1_struct_0(X6))|~m1_subset_1(X8,u1_struct_0(X6))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,X6))))|(~v2_pre_topc(X6)|(~v8_pre_topc(X6)|(~v3_orders_2(X6)|(~v4_orders_2(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|(~v2_lattice3(X6)|(~v1_compts_1(X6)|~l1_waybel_9(X6))))))))))&(~v5_pre_topc(k4_waybel_1(X6,esk1_4(X6,X7,X8,X9)),X6,X6)|(~v10_waybel_0(X7,X6)|~r3_waybel_9(X6,X7,X8))|X8!=X9|r2_lattice3(X6,k2_relset_1(u1_struct_0(X7),u1_struct_0(X6),u1_waybel_0(X6,X7)),X9)|~m1_subset_1(X9,u1_struct_0(X6))|~m1_subset_1(X8,u1_struct_0(X6))|(v2_struct_0(X7)|(~v4_orders_2(X7)|(~v7_waybel_0(X7)|~l1_waybel_0(X7,X6))))|(~v2_pre_topc(X6)|(~v8_pre_topc(X6)|(~v3_orders_2(X6)|(~v4_orders_2(X6)|(~v5_orders_2(X6)|(~v1_lattice3(X6)|(~v2_lattice3(X6)|(~v1_compts_1(X6)|~l1_waybel_9(X6))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l48_waybel_9])])])])])])).
fof(c_0_15, negated_conjecture, ~(![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((~(v2_struct_0(X3))&(v4_orders_2(X3)&(v7_waybel_0(X3)&l1_waybel_0(X3,X1))))=>((![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X4),X1,X1))&(v10_waybel_0(X3,X1)&r3_waybel_9(X1,X3,X2)))=>X2=k1_waybel_2(X1,X3)))))), inference(assume_negation,[status(cth)],[t35_waybel_9])).
cnf(c_0_16, plain, (m1_subset_1(esk1_4(X1,X2,X3,X4),u1_struct_0(X1))|r2_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X4)|v2_struct_0(X2)|~v10_waybel_0(X2,X1)|~r3_waybel_9(X1,X2,X3)|X3!=X4|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v1_compts_1(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_17, negated_conjecture, ![X14]:((v2_pre_topc(esk2_0)&(v8_pre_topc(esk2_0)&(v3_orders_2(esk2_0)&(v4_orders_2(esk2_0)&(v5_orders_2(esk2_0)&(v1_lattice3(esk2_0)&(v2_lattice3(esk2_0)&(v1_compts_1(esk2_0)&l1_waybel_9(esk2_0)))))))))&(m1_subset_1(esk3_0,u1_struct_0(esk2_0))&((~v2_struct_0(esk4_0)&(v4_orders_2(esk4_0)&(v7_waybel_0(esk4_0)&l1_waybel_0(esk4_0,esk2_0))))&(((~m1_subset_1(X14,u1_struct_0(esk2_0))|v5_pre_topc(k4_waybel_1(esk2_0,X14),esk2_0,esk2_0))&(v10_waybel_0(esk4_0,esk2_0)&r3_waybel_9(esk2_0,esk4_0,esk3_0)))&esk3_0!=k1_waybel_2(esk2_0,esk4_0))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_15])])])])])).
cnf(c_0_18, plain, (r2_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X3)|m1_subset_1(esk1_4(X1,X2,X3,X3),u1_struct_0(X1))|v2_struct_0(X2)|~r3_waybel_9(X1,X2,X3)|~v10_waybel_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_19, negated_conjecture, (v10_waybel_0(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20, negated_conjecture, (l1_waybel_0(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21, negated_conjecture, (v7_waybel_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, negated_conjecture, (l1_waybel_9(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, negated_conjecture, (v1_compts_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, negated_conjecture, (v2_lattice3(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, negated_conjecture, (v1_lattice3(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, negated_conjecture, (v5_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27, negated_conjecture, (v4_orders_2(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28, negated_conjecture, (v4_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_29, negated_conjecture, (v3_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_30, negated_conjecture, (v8_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_31, negated_conjecture, (v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_32, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_33, negated_conjecture, (r2_lattice3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),X1)|m1_subset_1(esk1_4(esk2_0,esk4_0,X1,X1),u1_struct_0(esk2_0))|~r3_waybel_9(esk2_0,esk4_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31])]), c_0_32])).
cnf(c_0_34, negated_conjecture, (r3_waybel_9(esk2_0,esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_35, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_36, plain, ![X15, X16, X17, X18, X20]:((m1_subset_1(esk5_4(X15,X16,X17,X18),u1_struct_0(X15))|~r3_waybel_9(X15,X16,X17)|X17!=X18|(~m1_subset_1(X20,u1_struct_0(X15))|(~r2_lattice3(X15,k2_relset_1(u1_struct_0(X16),u1_struct_0(X15),u1_waybel_0(X15,X16)),X20)|r3_orders_2(X15,X18,X20)))|~m1_subset_1(X18,u1_struct_0(X15))|~m1_subset_1(X17,u1_struct_0(X15))|(v2_struct_0(X16)|(~v4_orders_2(X16)|(~v7_waybel_0(X16)|~l1_waybel_0(X16,X15))))|(~v2_pre_topc(X15)|(~v8_pre_topc(X15)|(~v3_orders_2(X15)|(~v4_orders_2(X15)|(~v5_orders_2(X15)|(~v1_lattice3(X15)|(~v2_lattice3(X15)|(~v1_compts_1(X15)|~l1_waybel_9(X15))))))))))&(~v5_pre_topc(k4_waybel_1(X15,esk5_4(X15,X16,X17,X18)),X15,X15)|~r3_waybel_9(X15,X16,X17)|X17!=X18|(~m1_subset_1(X20,u1_struct_0(X15))|(~r2_lattice3(X15,k2_relset_1(u1_struct_0(X16),u1_struct_0(X15),u1_waybel_0(X15,X16)),X20)|r3_orders_2(X15,X18,X20)))|~m1_subset_1(X18,u1_struct_0(X15))|~m1_subset_1(X17,u1_struct_0(X15))|(v2_struct_0(X16)|(~v4_orders_2(X16)|(~v7_waybel_0(X16)|~l1_waybel_0(X16,X15))))|(~v2_pre_topc(X15)|(~v8_pre_topc(X15)|(~v3_orders_2(X15)|(~v4_orders_2(X15)|(~v5_orders_2(X15)|(~v1_lattice3(X15)|(~v2_lattice3(X15)|(~v1_compts_1(X15)|~l1_waybel_9(X15))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l49_waybel_9])])])])])])).
cnf(c_0_37, plain, (r2_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X4)|v2_struct_0(X2)|~v5_pre_topc(k4_waybel_1(X1,esk1_4(X1,X2,X3,X4)),X1,X1)|~v10_waybel_0(X2,X1)|~r3_waybel_9(X1,X2,X3)|X3!=X4|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v1_compts_1(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_38, negated_conjecture, (v5_pre_topc(k4_waybel_1(esk2_0,X1),esk2_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_39, negated_conjecture, (r2_lattice3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk3_0)|m1_subset_1(esk1_4(esk2_0,esk4_0,esk3_0,esk3_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])])).
fof(c_0_40, plain, ![X21]:((l1_pre_topc(X21)|~l1_waybel_9(X21))&(l1_orders_2(X21)|~l1_waybel_9(X21))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_waybel_9])])])).
cnf(c_0_41, plain, (m1_subset_1(esk5_4(X1,X2,X3,X4),u1_struct_0(X1))|r3_orders_2(X1,X4,X5)|v2_struct_0(X2)|~r3_waybel_9(X1,X2,X3)|X3!=X4|~m1_subset_1(X5,u1_struct_0(X1))|~r2_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X5)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v1_compts_1(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_42, plain, ![X22, X23, X25, X26, X27]:(((m1_subset_1(esk6_2(X22,X23),u1_struct_0(X22))|~r1_yellow_0(X22,X23)|(~v5_orders_2(X22)|~l1_orders_2(X22)))&((r2_lattice3(X22,X23,esk6_2(X22,X23))|~r1_yellow_0(X22,X23)|(~v5_orders_2(X22)|~l1_orders_2(X22)))&(~m1_subset_1(X25,u1_struct_0(X22))|(~r2_lattice3(X22,X23,X25)|r1_orders_2(X22,esk6_2(X22,X23),X25))|~r1_yellow_0(X22,X23)|(~v5_orders_2(X22)|~l1_orders_2(X22)))))&((m1_subset_1(esk7_3(X22,X26,X27),u1_struct_0(X22))|~r2_lattice3(X22,X26,X27)|~m1_subset_1(X27,u1_struct_0(X22))|r1_yellow_0(X22,X26)|(~v5_orders_2(X22)|~l1_orders_2(X22)))&((r2_lattice3(X22,X26,esk7_3(X22,X26,X27))|~r2_lattice3(X22,X26,X27)|~m1_subset_1(X27,u1_struct_0(X22))|r1_yellow_0(X22,X26)|(~v5_orders_2(X22)|~l1_orders_2(X22)))&(~r1_orders_2(X22,X27,esk7_3(X22,X26,X27))|~r2_lattice3(X22,X26,X27)|~m1_subset_1(X27,u1_struct_0(X22))|r1_yellow_0(X22,X26)|(~v5_orders_2(X22)|~l1_orders_2(X22)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_yellow_0])])])])])])).
cnf(c_0_43, plain, (r2_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X3)|v2_struct_0(X2)|~r3_waybel_9(X1,X2,X3)|~v10_waybel_0(X2,X1)|~v5_pre_topc(k4_waybel_1(X1,esk1_4(X1,X2,X3,X3)),X1,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_44, negated_conjecture, (r2_lattice3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk3_0)|v5_pre_topc(k4_waybel_1(esk2_0,esk1_4(esk2_0,esk4_0,esk3_0,esk3_0)),esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_45, plain, (l1_orders_2(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_46, plain, (r3_orders_2(X1,X2,X3)|m1_subset_1(esk5_4(X1,X4,X2,X2),u1_struct_0(X1))|v2_struct_0(X4)|~r2_lattice3(X1,k2_relset_1(u1_struct_0(X4),u1_struct_0(X1),u1_waybel_0(X1,X4)),X3)|~r3_waybel_9(X1,X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_waybel_0(X4,X1)|~v7_waybel_0(X4)|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X4)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_47, plain, (r2_lattice3(X1,X2,esk7_3(X1,X2,X3))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_48, negated_conjecture, (r2_lattice3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_34]), c_0_19]), c_0_35]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31])]), c_0_32])).
cnf(c_0_49, negated_conjecture, (l1_orders_2(esk2_0)), inference(spm,[status(thm)],[c_0_45, c_0_22])).
cnf(c_0_50, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_51, plain, ![X29]:(~l1_orders_2(X29)|(~v1_lattice3(X29)|~v2_struct_0(X29))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattice3])])])).
cnf(c_0_52, negated_conjecture, (r3_orders_2(esk2_0,esk3_0,X1)|m1_subset_1(esk5_4(esk2_0,esk4_0,esk3_0,esk3_0),u1_struct_0(esk2_0))|~r2_lattice3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_34]), c_0_35]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31])]), c_0_32])).
cnf(c_0_53, negated_conjecture, (r1_yellow_0(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)))|r2_lattice3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk7_3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49]), c_0_35]), c_0_26])])).
cnf(c_0_54, negated_conjecture, (r1_yellow_0(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)))|m1_subset_1(esk7_3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk3_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_48]), c_0_49]), c_0_35]), c_0_26])])).
cnf(c_0_55, plain, (~l1_orders_2(X1)|~v1_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_56, plain, (r3_orders_2(X1,X4,X5)|v2_struct_0(X2)|~v5_pre_topc(k4_waybel_1(X1,esk5_4(X1,X2,X3,X4)),X1,X1)|~r3_waybel_9(X1,X2,X3)|X3!=X4|~m1_subset_1(X5,u1_struct_0(X1))|~r2_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X5)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v1_compts_1(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_57, negated_conjecture, (r1_yellow_0(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)))|r3_orders_2(esk2_0,esk3_0,esk7_3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk3_0))|m1_subset_1(esk5_4(esk2_0,esk4_0,esk3_0,esk3_0),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54])).
fof(c_0_58, plain, ![X30, X31]:(v2_struct_0(X30)|~l1_orders_2(X30)|(~l1_waybel_0(X31,X30)|k1_waybel_2(X30,X31)=k4_yellow_2(X30,u1_waybel_0(X30,X31)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_waybel_2])])])])).
cnf(c_0_59, negated_conjecture, (~l1_orders_2(esk2_0)|~v2_struct_0(esk2_0)), inference(spm,[status(thm)],[c_0_55, c_0_25])).
fof(c_0_60, plain, ![X32, X33]:((v1_funct_1(u1_waybel_0(X32,X33))|(~l1_struct_0(X32)|~l1_waybel_0(X33,X32)))&((v1_funct_2(u1_waybel_0(X32,X33),u1_struct_0(X33),u1_struct_0(X32))|(~l1_struct_0(X32)|~l1_waybel_0(X33,X32)))&(m1_subset_1(u1_waybel_0(X32,X33),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X33),u1_struct_0(X32))))|(~l1_struct_0(X32)|~l1_waybel_0(X33,X32))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_waybel_0])])])).
cnf(c_0_61, plain, (r3_orders_2(X1,X2,X3)|v2_struct_0(X4)|~r2_lattice3(X1,k2_relset_1(u1_struct_0(X4),u1_struct_0(X1),u1_waybel_0(X1,X4)),X3)|~r3_waybel_9(X1,X4,X2)|~v5_pre_topc(k4_waybel_1(X1,esk5_4(X1,X4,X2,X2)),X1,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_waybel_0(X4,X1)|~v7_waybel_0(X4)|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X4)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_56])).
cnf(c_0_62, negated_conjecture, (r1_yellow_0(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)))|r3_orders_2(esk2_0,esk3_0,esk7_3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk3_0))|v5_pre_topc(k4_waybel_1(esk2_0,esk5_4(esk2_0,esk4_0,esk3_0,esk3_0)),esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_38, c_0_57])).
cnf(c_0_63, plain, (v2_struct_0(X1)|k1_waybel_2(X1,X2)=k4_yellow_2(X1,u1_waybel_0(X1,X2))|~l1_orders_2(X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_64, negated_conjecture, (~v2_struct_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_49])])).
cnf(c_0_65, plain, (m1_subset_1(u1_waybel_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
fof(c_0_66, plain, ![X34]:(~l1_orders_2(X34)|l1_struct_0(X34)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
fof(c_0_67, plain, ![X35, X36, X37]:((~r3_orders_2(X35,X36,X37)|r1_orders_2(X35,X36,X37)|(v2_struct_0(X35)|(~v3_orders_2(X35)|~l1_orders_2(X35))|(~m1_subset_1(X36,u1_struct_0(X35))|~m1_subset_1(X37,u1_struct_0(X35)))))&(~r1_orders_2(X35,X36,X37)|r3_orders_2(X35,X36,X37)|(v2_struct_0(X35)|(~v3_orders_2(X35)|~l1_orders_2(X35))|(~m1_subset_1(X36,u1_struct_0(X35))|~m1_subset_1(X37,u1_struct_0(X35)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_68, negated_conjecture, (r1_yellow_0(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)))|r3_orders_2(esk2_0,esk3_0,esk7_3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk3_0))|r3_orders_2(esk2_0,esk3_0,X1)|~r2_lattice3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_34]), c_0_35]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31])]), c_0_32])).
cnf(c_0_69, negated_conjecture, (esk3_0!=k1_waybel_2(esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_70, negated_conjecture, (k1_waybel_2(esk2_0,esk4_0)=k4_yellow_2(esk2_0,u1_waybel_0(esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_20]), c_0_49])]), c_0_64])).
fof(c_0_71, plain, ![X38, X39]:(v2_struct_0(X38)|~l1_orders_2(X38)|(~v1_relat_1(X39)|k4_yellow_2(X38,X39)=k1_yellow_0(X38,k2_relat_1(X39)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_yellow_2])])])])).
fof(c_0_72, plain, ![X40, X41, X42]:(~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(X40,X41)))|v1_relat_1(X42)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_73, negated_conjecture, (m1_subset_1(u1_waybel_0(esk2_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~l1_struct_0(esk2_0)), inference(spm,[status(thm)],[c_0_65, c_0_20])).
cnf(c_0_74, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_75, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_76, negated_conjecture, (r1_yellow_0(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)))|r3_orders_2(esk2_0,esk3_0,esk7_3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk3_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_53]), c_0_54])).
cnf(c_0_77, negated_conjecture, (k4_yellow_2(esk2_0,u1_waybel_0(esk2_0,esk4_0))!=esk3_0), inference(rw,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_78, plain, (v2_struct_0(X1)|k4_yellow_2(X1,X2)=k1_yellow_0(X1,k2_relat_1(X2))|~l1_orders_2(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_79, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_80, negated_conjecture, (m1_subset_1(u1_waybel_0(esk2_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_49])])).
fof(c_0_81, plain, ![X43, X44, X45]:(~m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(X43,X44)))|k2_relset_1(X43,X44,X45)=k2_relat_1(X45)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
fof(c_0_82, plain, ![X46, X47, X48, X49]:(((r2_lattice3(X46,X47,X48)|X48!=k1_yellow_0(X46,X47)|~r1_yellow_0(X46,X47)|~m1_subset_1(X48,u1_struct_0(X46))|~l1_orders_2(X46))&(~m1_subset_1(X49,u1_struct_0(X46))|(~r2_lattice3(X46,X47,X49)|r1_orders_2(X46,X48,X49))|X48!=k1_yellow_0(X46,X47)|~r1_yellow_0(X46,X47)|~m1_subset_1(X48,u1_struct_0(X46))|~l1_orders_2(X46)))&((m1_subset_1(esk8_3(X46,X47,X48),u1_struct_0(X46))|~r2_lattice3(X46,X47,X48)|X48=k1_yellow_0(X46,X47)|~r1_yellow_0(X46,X47)|~m1_subset_1(X48,u1_struct_0(X46))|~l1_orders_2(X46))&((r2_lattice3(X46,X47,esk8_3(X46,X47,X48))|~r2_lattice3(X46,X47,X48)|X48=k1_yellow_0(X46,X47)|~r1_yellow_0(X46,X47)|~m1_subset_1(X48,u1_struct_0(X46))|~l1_orders_2(X46))&(~r1_orders_2(X46,X48,esk8_3(X46,X47,X48))|~r2_lattice3(X46,X47,X48)|X48=k1_yellow_0(X46,X47)|~r1_yellow_0(X46,X47)|~m1_subset_1(X48,u1_struct_0(X46))|~l1_orders_2(X46))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
cnf(c_0_83, plain, (r1_yellow_0(X1,X3)|~r1_orders_2(X1,X2,esk7_3(X1,X3,X2))|~r2_lattice3(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_84, negated_conjecture, (r1_orders_2(esk2_0,esk3_0,esk7_3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk3_0))|r1_yellow_0(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_49]), c_0_35]), c_0_29])]), c_0_64]), c_0_54])).
cnf(c_0_85, negated_conjecture, (k1_yellow_0(esk2_0,k2_relat_1(u1_waybel_0(esk2_0,esk4_0)))!=esk3_0|~v1_relat_1(u1_waybel_0(esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_49])]), c_0_64])).
cnf(c_0_86, negated_conjecture, (v1_relat_1(u1_waybel_0(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_87, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_88, plain, (r2_lattice3(X1,X2,esk8_3(X1,X2,X3))|X3=k1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_89, negated_conjecture, (r1_yellow_0(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_49]), c_0_48]), c_0_35]), c_0_26])])).
cnf(c_0_90, negated_conjecture, (k1_yellow_0(esk2_0,k2_relat_1(u1_waybel_0(esk2_0,esk4_0)))!=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_86])])).
cnf(c_0_91, negated_conjecture, (k2_relat_1(u1_waybel_0(esk2_0,esk4_0))=k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_87, c_0_80])).
cnf(c_0_92, plain, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|X3=k1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_93, negated_conjecture, (X1=k1_yellow_0(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)))|r2_lattice3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk8_3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),X1))|~r2_lattice3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_49])])).
cnf(c_0_94, negated_conjecture, (esk3_0!=k1_yellow_0(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)))), inference(rw,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_95, negated_conjecture, (X1=k1_yellow_0(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)))|m1_subset_1(esk8_3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),X1),u1_struct_0(esk2_0))|~r2_lattice3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_89]), c_0_49])])).
cnf(c_0_96, negated_conjecture, (r2_lattice3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk8_3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_48]), c_0_35])]), c_0_94])).
cnf(c_0_97, negated_conjecture, (m1_subset_1(esk8_3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_48]), c_0_35])]), c_0_94])).
cnf(c_0_98, negated_conjecture, (r3_orders_2(esk2_0,esk3_0,esk8_3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk3_0))|m1_subset_1(esk5_4(esk2_0,esk4_0,esk3_0,esk3_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_96]), c_0_97])])).
cnf(c_0_99, negated_conjecture, (r3_orders_2(esk2_0,esk3_0,esk8_3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk3_0))|v5_pre_topc(k4_waybel_1(esk2_0,esk5_4(esk2_0,esk4_0,esk3_0,esk3_0)),esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_38, c_0_98])).
cnf(c_0_100, negated_conjecture, (r3_orders_2(esk2_0,esk3_0,esk8_3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk3_0))|r3_orders_2(esk2_0,esk3_0,X1)|~r2_lattice3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_99]), c_0_34]), c_0_35]), c_0_20]), c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31])]), c_0_32])).
cnf(c_0_101, negated_conjecture, (r3_orders_2(esk2_0,esk3_0,esk8_3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_96]), c_0_97])])).
cnf(c_0_102, plain, (X2=k1_yellow_0(X1,X3)|~r1_orders_2(X1,X2,esk8_3(X1,X3,X2))|~r2_lattice3(X1,X3,X2)|~r1_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_103, negated_conjecture, (r1_orders_2(esk2_0,esk3_0,esk8_3(esk2_0,k2_relset_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0),u1_waybel_0(esk2_0,esk4_0)),esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_101]), c_0_49]), c_0_97]), c_0_35]), c_0_29])]), c_0_64])).
cnf(c_0_104, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_89]), c_0_49]), c_0_48]), c_0_35])]), c_0_94]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 105
# Proof object clause steps            : 76
# Proof object formula steps           : 29
# Proof object conjectures             : 56
# Proof object clause conjectures      : 53
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 37
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 31
# Proof object simplifying inferences  : 145
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 46
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 46
# Processed clauses                    : 164
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 164
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 23
# Generated clauses                    : 97
# ...of the previous two non-trivial   : 85
# Contextual simplify-reflections      : 3
# Paramodulations                      : 91
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 86
#    Positive orientable unit clauses  : 41
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 41
# Current number of unprocessed clauses: 10
# ...number of literals in the above   : 30
# Current number of archived formulas  : 0
# Current number of archived clauses   : 72
# Clause-clause subsumption calls (NU) : 4168
# Rec. Clause-clause subsumption calls : 483
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 50
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 34
# BW rewrite match successes           : 12
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 9371

# -------------------------------------------------
# User time                : 0.014 s
# System time              : 0.000 s
# Total time               : 0.014 s
# Maximum resident set size: 28740 pages
