# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t36_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t100_xboole_1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t70_enumset1)).
fof(t43_zfmisc_1, conjecture, ![X1, X2, X3]:~((k1_tarski(X1)=k2_xboole_0(X2,X3)&(~((X2=k1_tarski(X1)&X3=k1_tarski(X1)))&(~((X2=k1_xboole_0&X3=k1_tarski(X1)))&~((X2=k1_tarski(X1)&X3=k1_xboole_0)))))), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t43_zfmisc_1)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t2_boole)).
fof(t5_boole, axiom, ![X1]:k5_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t5_boole)).
fof(t3_xboole_1, axiom, ![X1]:(r1_tarski(X1,k1_xboole_0)=>X1=k1_xboole_0), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t3_xboole_1)).
fof(t108_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(k3_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t108_xboole_1)).
fof(t1_xboole_0, axiom, ![X1, X2, X3]:(r2_hidden(X1,k5_xboole_0(X2,X3))<=>~((r2_hidden(X1,X2)<=>r2_hidden(X1,X3)))), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t1_xboole_0)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t7_xboole_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t75_enumset1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t69_enumset1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t4_xboole_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', d1_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t28_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', commutativity_k3_xboole_0)).
fof(l31_zfmisc_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>k3_xboole_0(X2,k1_tarski(X1))=k1_tarski(X1)), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', l31_zfmisc_1)).
fof(l27_zfmisc_1, axiom, ![X1, X2]:(~(r2_hidden(X1,X2))=>r1_xboole_0(k1_tarski(X1),X2)), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', l27_zfmisc_1)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', l13_xboole_0)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', l32_xboole_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t12_xboole_1)).
fof(t38_zfmisc_1, axiom, ![X1, X2, X3]:(r1_tarski(k2_tarski(X1,X2),X3)<=>(r2_hidden(X1,X3)&r2_hidden(X2,X3))), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t38_zfmisc_1)).
fof(t10_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(X1,k2_xboole_0(X3,X2))), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t10_xboole_1)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/4aed1b26f2314a0dbecae8e3bd0b362d.E_input', t83_xboole_1)).
fof(c_0_29, plain, ![X52, X53]:r1_tarski(k4_xboole_0(X52,X53),X52), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_30, plain, ![X47, X48]:k4_xboole_0(X47,X48)=k5_xboole_0(X47,k3_xboole_0(X47,X48)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_31, plain, ![X58, X59]:k3_tarski(k2_tarski(X58,X59))=k2_xboole_0(X58,X59), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_32, plain, ![X56, X57]:k1_enumset1(X56,X56,X57)=k2_tarski(X56,X57), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_33, negated_conjecture, ~(![X1, X2, X3]:~((k1_tarski(X1)=k2_xboole_0(X2,X3)&(~((X2=k1_tarski(X1)&X3=k1_tarski(X1)))&(~((X2=k1_xboole_0&X3=k1_tarski(X1)))&~((X2=k1_tarski(X1)&X3=k1_xboole_0))))))), inference(assume_negation,[status(cth)],[t43_zfmisc_1])).
cnf(c_0_34, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_35, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_36, plain, ![X41]:k3_xboole_0(X41,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_37, plain, ![X25]:k5_xboole_0(X25,k1_xboole_0)=X25, inference(variable_rename,[status(thm)],[t5_boole])).
fof(c_0_38, plain, ![X60]:(~r1_tarski(X60,k1_xboole_0)|X60=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_xboole_1])])).
fof(c_0_39, plain, ![X20, X21, X22]:(~r1_tarski(X20,X21)|r1_tarski(k3_xboole_0(X20,X22),X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t108_xboole_1])])).
fof(c_0_40, plain, ![X42, X43, X44]:(((~r2_hidden(X42,X43)|~r2_hidden(X42,X44)|~r2_hidden(X42,k5_xboole_0(X43,X44)))&(r2_hidden(X42,X43)|r2_hidden(X42,X44)|~r2_hidden(X42,k5_xboole_0(X43,X44))))&((~r2_hidden(X42,X43)|r2_hidden(X42,X44)|r2_hidden(X42,k5_xboole_0(X43,X44)))&(~r2_hidden(X42,X44)|r2_hidden(X42,X43)|r2_hidden(X42,k5_xboole_0(X43,X44))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_0])])])).
fof(c_0_41, plain, ![X23, X24]:r1_tarski(X23,k2_xboole_0(X23,X24)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_42, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_43, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_44, plain, ![X68, X69, X70]:k2_enumset1(X68,X68,X69,X70)=k1_enumset1(X68,X69,X70), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_45, plain, ![X11, X12, X13, X14]:k3_enumset1(X11,X11,X12,X13,X14)=k2_enumset1(X11,X12,X13,X14), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_46, plain, ![X81, X82, X83, X84, X85]:k4_enumset1(X81,X81,X82,X83,X84,X85)=k3_enumset1(X81,X82,X83,X84,X85), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_47, plain, ![X75, X76, X77, X78, X79, X80]:k5_enumset1(X75,X75,X76,X77,X78,X79,X80)=k4_enumset1(X75,X76,X77,X78,X79,X80), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_48, plain, ![X26, X27, X28, X29, X30, X31, X32]:k6_enumset1(X26,X26,X27,X28,X29,X30,X31,X32)=k5_enumset1(X26,X27,X28,X29,X30,X31,X32), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_49, negated_conjecture, (k1_tarski(esk1_0)=k2_xboole_0(esk2_0,esk3_0)&((esk2_0!=k1_tarski(esk1_0)|esk3_0!=k1_tarski(esk1_0))&((esk2_0!=k1_xboole_0|esk3_0!=k1_tarski(esk1_0))&(esk2_0!=k1_tarski(esk1_0)|esk3_0!=k1_xboole_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_33])])])).
fof(c_0_50, plain, ![X19]:k2_tarski(X19,X19)=k1_tarski(X19), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_51, plain, ![X33, X34, X36, X37, X38]:((r1_xboole_0(X33,X34)|r2_hidden(esk5_2(X33,X34),k3_xboole_0(X33,X34)))&(~r2_hidden(X38,k3_xboole_0(X36,X37))|~r1_xboole_0(X36,X37))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_52, plain, ![X15, X16, X17]:((~v1_xboole_0(X15)|~r2_hidden(X16,X15))&(r2_hidden(esk4_1(X17),X17)|v1_xboole_0(X17))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_53, plain, ![X45, X46]:(~r1_tarski(X45,X46)|k3_xboole_0(X45,X46)=X45), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_54, plain, (r1_tarski(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X1)), inference(rw,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_55, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_56, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_57, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_58, plain, (r1_tarski(k3_xboole_0(X1,X3),X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_59, plain, (r2_hidden(X1,X3)|r2_hidden(X1,k5_xboole_0(X3,X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_60, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_61, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_62, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_63, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_64, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_65, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_66, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_67, negated_conjecture, (k1_tarski(esk1_0)=k2_xboole_0(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_68, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
fof(c_0_69, plain, ![X71, X72]:k3_xboole_0(X71,X72)=k3_xboole_0(X72,X71), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_70, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_71, plain, (r2_hidden(esk4_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_72, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_73, plain, (r1_tarski(X1,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56])).
cnf(c_0_74, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_75, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r2_hidden(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_76, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_59, c_0_56])).
fof(c_0_77, plain, ![X66, X67]:(~r2_hidden(X66,X67)|k3_xboole_0(X67,k1_tarski(X66))=k1_tarski(X66)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l31_zfmisc_1])])).
cnf(c_0_78, plain, (r1_tarski(X1,k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_61]), c_0_62]), c_0_63]), c_0_64]), c_0_65]), c_0_66])).
cnf(c_0_79, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=k3_tarski(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_68]), c_0_43]), c_0_61]), c_0_62]), c_0_62]), c_0_63]), c_0_63]), c_0_64]), c_0_64]), c_0_65]), c_0_65]), c_0_66]), c_0_66])).
cnf(c_0_80, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_81, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_82, plain, (k3_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_83, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk5_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_84, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_74, c_0_58])).
cnf(c_0_85, plain, (~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_56]), c_0_76])).
fof(c_0_86, plain, ![X54, X55]:(r2_hidden(X54,X55)|r1_xboole_0(k1_tarski(X54),X55)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l27_zfmisc_1])])])).
cnf(c_0_87, plain, (k3_xboole_0(X2,k1_tarski(X1))=k1_tarski(X1)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_88, negated_conjecture, (r1_tarski(esk2_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))), inference(spm,[status(thm)],[c_0_78, c_0_79])).
fof(c_0_89, plain, ![X86]:(~v1_xboole_0(X86)|X86=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_90, plain, (~r1_xboole_0(X1,X2)|~r2_hidden(X3,k3_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_70, c_0_80])).
cnf(c_0_91, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_92, plain, (v1_xboole_0(X1)|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_93, plain, (r1_xboole_0(k3_xboole_0(X1,X2),X3)|~r1_tarski(X1,k1_xboole_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_85])).
fof(c_0_94, plain, ![X61, X62]:((k4_xboole_0(X61,X62)!=k1_xboole_0|r1_tarski(X61,X62))&(~r1_tarski(X61,X62)|k4_xboole_0(X61,X62)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
fof(c_0_95, plain, ![X39, X40]:(~r1_tarski(X39,X40)|k2_xboole_0(X39,X40)=X40), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_96, plain, (r2_hidden(X1,X2)|r1_xboole_0(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_97, plain, (k3_xboole_0(X2,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))=k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_68]), c_0_68]), c_0_43]), c_0_43]), c_0_62]), c_0_62]), c_0_63]), c_0_63]), c_0_64]), c_0_64]), c_0_65]), c_0_65]), c_0_66]), c_0_66])).
cnf(c_0_98, negated_conjecture, (k3_xboole_0(esk2_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))=esk2_0), inference(spm,[status(thm)],[c_0_72, c_0_88])).
cnf(c_0_99, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_100, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_90, c_0_71])).
cnf(c_0_101, plain, (r2_hidden(esk4_1(k5_xboole_0(X1,X2)),X2)|r2_hidden(esk4_1(k5_xboole_0(X1,X2)),X1)|v1_xboole_0(k5_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_91, c_0_71])).
cnf(c_0_102, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_103, plain, (k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2)))=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_54]), c_0_80])).
cnf(c_0_104, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_105, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_106, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_107, negated_conjecture, (esk2_0!=k1_tarski(esk1_0)|esk3_0!=k1_tarski(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_108, plain, ![X63, X64, X65]:(((r2_hidden(X63,X65)|~r1_tarski(k2_tarski(X63,X64),X65))&(r2_hidden(X64,X65)|~r1_tarski(k2_tarski(X63,X64),X65)))&(~r2_hidden(X63,X65)|~r2_hidden(X64,X65)|r1_tarski(k2_tarski(X63,X64),X65))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t38_zfmisc_1])])])).
cnf(c_0_109, plain, (r2_hidden(X1,X2)|r1_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_96, c_0_68]), c_0_43]), c_0_62]), c_0_63]), c_0_64]), c_0_65]), c_0_66])).
cnf(c_0_110, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=esk2_0|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_111, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_112, plain, (r2_hidden(esk4_1(k5_xboole_0(k1_xboole_0,X1)),X1)|v1_xboole_0(k5_xboole_0(k1_xboole_0,X1))), inference(spm,[status(thm)],[c_0_85, c_0_101])).
cnf(c_0_113, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_90, c_0_83])).
cnf(c_0_114, plain, (v1_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)))|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_115, plain, (r1_tarski(X1,X2)|k5_xboole_0(X1,k3_xboole_0(X1,X2))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_104, c_0_35])).
cnf(c_0_116, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=k1_xboole_0|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_105, c_0_35])).
cnf(c_0_117, plain, (k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))=X2|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_106, c_0_61]), c_0_62]), c_0_63]), c_0_64]), c_0_65]), c_0_66])).
cnf(c_0_118, negated_conjecture, (esk2_0!=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)|esk3_0!=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_107, c_0_68]), c_0_68]), c_0_43]), c_0_43]), c_0_62]), c_0_62]), c_0_63]), c_0_63]), c_0_64]), c_0_64]), c_0_65]), c_0_65]), c_0_66]), c_0_66])).
cnf(c_0_119, negated_conjecture, (esk2_0!=k1_xboole_0|esk3_0!=k1_tarski(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_120, plain, (r1_tarski(k3_xboole_0(X1,X2),X3)|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_58, c_0_80])).
cnf(c_0_121, plain, (r1_tarski(k2_tarski(X1,X3),X2)|~r2_hidden(X1,X2)|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_108])).
cnf(c_0_122, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_99, c_0_81])).
cnf(c_0_123, negated_conjecture, (r1_xboole_0(esk2_0,X1)|r2_hidden(esk1_0,X1)|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_109, c_0_110])).
cnf(c_0_124, plain, (k3_xboole_0(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2))=k1_xboole_0|r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_111, c_0_109])).
cnf(c_0_125, plain, (v1_xboole_0(k5_xboole_0(k1_xboole_0,k3_xboole_0(X1,X2)))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_70, c_0_112])).
cnf(c_0_126, negated_conjecture, (~r1_xboole_0(esk2_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_70, c_0_98])).
cnf(c_0_127, plain, (r1_xboole_0(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2))|r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_113, c_0_109])).
cnf(c_0_128, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_99, c_0_114])).
cnf(c_0_129, plain, (r1_tarski(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2)))|k5_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2)))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_115, c_0_103])).
cnf(c_0_130, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_82]), c_0_73])])).
cnf(c_0_131, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=esk3_0|~r1_tarski(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_79, c_0_117])).
cnf(c_0_132, negated_conjecture, (esk3_0!=esk2_0|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_118, c_0_110])).
cnf(c_0_133, negated_conjecture, (esk2_0!=k1_xboole_0|esk3_0!=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_119, c_0_68]), c_0_43]), c_0_62]), c_0_63]), c_0_64]), c_0_65]), c_0_66])).
cnf(c_0_134, negated_conjecture, (r1_tarski(esk2_0,X1)|~r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),X1)), inference(spm,[status(thm)],[c_0_120, c_0_98])).
cnf(c_0_135, plain, (r1_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X3),X2)|~r2_hidden(X3,X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_121, c_0_43]), c_0_62]), c_0_63]), c_0_64]), c_0_65]), c_0_66])).
cnf(c_0_136, negated_conjecture, (k3_xboole_0(esk2_0,X1)=k1_xboole_0|r2_hidden(esk1_0,X1)|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_137, negated_conjecture, (esk2_0=k1_xboole_0|r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_98, c_0_124])).
cnf(c_0_138, negated_conjecture, (esk2_0!=k1_tarski(esk1_0)|esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_139, plain, (k5_xboole_0(k1_xboole_0,k3_xboole_0(X1,X2))=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_99, c_0_125])).
cnf(c_0_140, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_141, plain, (r1_xboole_0(X1,X1)|r2_hidden(esk5_2(X1,X1),X1)), inference(spm,[status(thm)],[c_0_83, c_0_82])).
cnf(c_0_142, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_115, c_0_128])).
cnf(c_0_143, negated_conjecture, (r1_tarski(esk2_0,k1_xboole_0)|esk2_0!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_98]), c_0_130]), c_0_130]), c_0_56])).
cnf(c_0_144, negated_conjecture, (~r1_tarski(esk2_0,esk3_0)|~r2_hidden(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_110]), c_0_132])).
cnf(c_0_145, negated_conjecture, (esk2_0!=k1_xboole_0|~r1_tarski(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_133, c_0_131])).
cnf(c_0_146, negated_conjecture, (r1_tarski(esk2_0,X1)|~r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_134, c_0_135])).
cnf(c_0_147, negated_conjecture, (k3_xboole_0(esk2_0,X1)=k1_xboole_0|esk2_0=k1_xboole_0|r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_136, c_0_137])).
cnf(c_0_148, negated_conjecture, (esk3_0!=k1_xboole_0|esk2_0!=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_138, c_0_68]), c_0_43]), c_0_62]), c_0_63]), c_0_64]), c_0_65]), c_0_66])).
cnf(c_0_149, plain, (k5_xboole_0(k1_xboole_0,X1)=k1_xboole_0|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_139, c_0_82])).
cnf(c_0_150, negated_conjecture, (r1_xboole_0(esk2_0,esk2_0)|r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_140, c_0_141])).
fof(c_0_151, plain, ![X49, X50, X51]:(~r1_tarski(X49,X50)|r1_tarski(X49,k2_xboole_0(X51,X50))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_xboole_1])])).
cnf(c_0_152, negated_conjecture, (r1_tarski(esk2_0,X1)|esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_142, c_0_143])).
fof(c_0_153, plain, ![X73, X74]:((~r1_xboole_0(X73,X74)|k4_xboole_0(X73,X74)=X73)&(k4_xboole_0(X73,X74)!=X73|r1_xboole_0(X73,X74))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
cnf(c_0_154, negated_conjecture, (~r1_tarski(esk2_0,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_137]), c_0_145])).
cnf(c_0_155, negated_conjecture, (k3_xboole_0(esk2_0,X1)=k1_xboole_0|esk2_0=k1_xboole_0|r1_tarski(esk2_0,X1)), inference(spm,[status(thm)],[c_0_146, c_0_147])).
cnf(c_0_156, negated_conjecture, (esk3_0!=k1_xboole_0|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_148, c_0_110])).
cnf(c_0_157, negated_conjecture, (k5_xboole_0(k1_xboole_0,esk2_0)=k1_xboole_0|r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_149, c_0_150])).
cnf(c_0_158, plain, (r1_tarski(X1,k2_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_151])).
cnf(c_0_159, negated_conjecture, (esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_145, c_0_152])).
cnf(c_0_160, plain, (k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_161, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk5_2(X1,X2),k3_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_83, c_0_80])).
cnf(c_0_162, negated_conjecture, (k3_xboole_0(esk2_0,esk3_0)=k1_xboole_0|esk2_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_154, c_0_155])).
cnf(c_0_163, negated_conjecture, (k5_xboole_0(k1_xboole_0,esk2_0)=k1_xboole_0|esk3_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_156, c_0_157])).
cnf(c_0_164, plain, (r1_tarski(X1,k3_tarski(k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X2)))|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_158, c_0_61]), c_0_62]), c_0_63]), c_0_64]), c_0_65]), c_0_66])).
cnf(c_0_165, negated_conjecture, (r2_hidden(esk1_0,esk2_0)), inference(sr,[status(thm)],[c_0_137, c_0_159])).
cnf(c_0_166, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=X1|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_160, c_0_35])).
cnf(c_0_167, negated_conjecture, (esk2_0=k1_xboole_0|r1_xboole_0(esk3_0,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_162]), c_0_85])).
cnf(c_0_168, negated_conjecture, (esk3_0!=k1_xboole_0|~r2_hidden(X1,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_163]), c_0_85])).
cnf(c_0_169, negated_conjecture, (r1_tarski(X1,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_164, c_0_79])).
cnf(c_0_170, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_110, c_0_165])])).
cnf(c_0_171, plain, (X1=k1_xboole_0|~r1_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_116, c_0_166])).
cnf(c_0_172, negated_conjecture, (r1_xboole_0(esk3_0,esk2_0)), inference(sr,[status(thm)],[c_0_167, c_0_159])).
cnf(c_0_173, negated_conjecture, (esk3_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_168, c_0_165])).
cnf(c_0_174, negated_conjecture, (r1_tarski(X1,esk2_0)|~r1_tarski(X1,esk3_0)), inference(rw,[status(thm)],[c_0_169, c_0_170])).
cnf(c_0_175, negated_conjecture, (~r1_tarski(esk3_0,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_172]), c_0_173])).
cnf(c_0_176, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_73]), c_0_175]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 177
# Proof object clause steps            : 118
# Proof object formula steps           : 59
# Proof object conjectures             : 46
# Proof object clause conjectures      : 43
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 36
# Proof object initial formulas used   : 29
# Proof object generating inferences   : 63
# Proof object simplifying inferences  : 111
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 29
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 41
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 32
# Processed clauses                    : 2484
# ...of these trivial                  : 39
# ...subsumed                          : 1938
# ...remaining for further processing  : 507
# Other redundant clauses eliminated   : 46
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 44
# Backward-rewritten                   : 103
# Generated clauses                    : 13038
# ...of the previous two non-trivial   : 7874
# Contextual simplify-reflections      : 23
# Paramodulations                      : 12964
# Factorizations                       : 8
# Equation resolutions                 : 46
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 340
#    Positive orientable unit clauses  : 38
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 12
#    Non-unit-clauses                  : 289
# Current number of unprocessed clauses: 5340
# ...number of literals in the above   : 18187
# Current number of archived formulas  : 0
# Current number of archived clauses   : 176
# Clause-clause subsumption calls (NU) : 26532
# Rec. Clause-clause subsumption calls : 19473
# Non-unit clause-clause subsumptions  : 1643
# Unit Clause-clause subsumption calls : 1168
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 56
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 130630

# -------------------------------------------------
# User time                : 0.085 s
# System time              : 0.003 s
# Total time               : 0.089 s
# Maximum resident set size: 28840 pages
