# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d8_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(X3=k5_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>?[X6]:(r2_hidden(k4_tarski(X4,X6),X1)&r2_hidden(k4_tarski(X6,X5),X2))))))), file('data/MPTP/proofs/4afc037ae0b94cc7bad5028893eb39ac.E_input', d8_relat_1)).
fof(dt_k5_relat_1, axiom, ![X1, X2]:((v1_relat_1(X1)&v1_relat_1(X2))=>v1_relat_1(k5_relat_1(X1,X2))), file('data/MPTP/proofs/4afc037ae0b94cc7bad5028893eb39ac.E_input', dt_k5_relat_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/4afc037ae0b94cc7bad5028893eb39ac.E_input', t3_xboole_0)).
fof(t56_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>(![X2, X3]:~(r2_hidden(k4_tarski(X2,X3),X1))=>X1=k1_xboole_0)), file('data/MPTP/proofs/4afc037ae0b94cc7bad5028893eb39ac.E_input', t56_relat_1)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/4afc037ae0b94cc7bad5028893eb39ac.E_input', t83_xboole_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/4afc037ae0b94cc7bad5028893eb39ac.E_input', l32_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/4afc037ae0b94cc7bad5028893eb39ac.E_input', d10_xboole_0)).
fof(fc1_relat_1, axiom, ![X1, X2]:(v1_relat_1(X1)=>v1_relat_1(k3_xboole_0(X1,X2))), file('data/MPTP/proofs/4afc037ae0b94cc7bad5028893eb39ac.E_input', fc1_relat_1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/4afc037ae0b94cc7bad5028893eb39ac.E_input', t12_setfam_1)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/4afc037ae0b94cc7bad5028893eb39ac.E_input', t2_boole)).
fof(t62_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>(k5_relat_1(k1_xboole_0,X1)=k1_xboole_0&k5_relat_1(X1,k1_xboole_0)=k1_xboole_0)), file('data/MPTP/proofs/4afc037ae0b94cc7bad5028893eb39ac.E_input', t62_relat_1)).
fof(c_0_11, plain, ![X19, X20, X21, X22, X23, X25, X26, X27, X30]:((((r2_hidden(k4_tarski(X22,esk2_5(X19,X20,X21,X22,X23)),X19)|~r2_hidden(k4_tarski(X22,X23),X21)|X21!=k5_relat_1(X19,X20)|~v1_relat_1(X21)|~v1_relat_1(X20)|~v1_relat_1(X19))&(r2_hidden(k4_tarski(esk2_5(X19,X20,X21,X22,X23),X23),X20)|~r2_hidden(k4_tarski(X22,X23),X21)|X21!=k5_relat_1(X19,X20)|~v1_relat_1(X21)|~v1_relat_1(X20)|~v1_relat_1(X19)))&(~r2_hidden(k4_tarski(X25,X27),X19)|~r2_hidden(k4_tarski(X27,X26),X20)|r2_hidden(k4_tarski(X25,X26),X21)|X21!=k5_relat_1(X19,X20)|~v1_relat_1(X21)|~v1_relat_1(X20)|~v1_relat_1(X19)))&((~r2_hidden(k4_tarski(esk3_3(X19,X20,X21),esk4_3(X19,X20,X21)),X21)|(~r2_hidden(k4_tarski(esk3_3(X19,X20,X21),X30),X19)|~r2_hidden(k4_tarski(X30,esk4_3(X19,X20,X21)),X20))|X21=k5_relat_1(X19,X20)|~v1_relat_1(X21)|~v1_relat_1(X20)|~v1_relat_1(X19))&((r2_hidden(k4_tarski(esk3_3(X19,X20,X21),esk5_3(X19,X20,X21)),X19)|r2_hidden(k4_tarski(esk3_3(X19,X20,X21),esk4_3(X19,X20,X21)),X21)|X21=k5_relat_1(X19,X20)|~v1_relat_1(X21)|~v1_relat_1(X20)|~v1_relat_1(X19))&(r2_hidden(k4_tarski(esk5_3(X19,X20,X21),esk4_3(X19,X20,X21)),X20)|r2_hidden(k4_tarski(esk3_3(X19,X20,X21),esk4_3(X19,X20,X21)),X21)|X21=k5_relat_1(X19,X20)|~v1_relat_1(X21)|~v1_relat_1(X20)|~v1_relat_1(X19))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_relat_1])])])])])])).
fof(c_0_12, plain, ![X32, X33]:(~v1_relat_1(X32)|~v1_relat_1(X33)|v1_relat_1(k5_relat_1(X32,X33))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_relat_1])])).
fof(c_0_13, plain, ![X11, X12, X14, X15, X16]:(((r2_hidden(esk1_2(X11,X12),X11)|r1_xboole_0(X11,X12))&(r2_hidden(esk1_2(X11,X12),X12)|r1_xboole_0(X11,X12)))&(~r2_hidden(X16,X14)|~r2_hidden(X16,X15)|~r1_xboole_0(X14,X15))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_14, plain, (r2_hidden(k4_tarski(X1,esk2_5(X2,X3,X4,X1,X5)),X2)|~r2_hidden(k4_tarski(X1,X5),X4)|X4!=k5_relat_1(X2,X3)|~v1_relat_1(X4)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (v1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, plain, (r2_hidden(k4_tarski(X1,esk2_5(X2,X3,k5_relat_1(X2,X3),X1,X4)),X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~r2_hidden(k4_tarski(X1,X4),k5_relat_1(X2,X3))), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_14]), c_0_15])).
fof(c_0_18, plain, ![X34]:(~v1_relat_1(X34)|(r2_hidden(k4_tarski(esk6_1(X34),esk7_1(X34)),X34)|X34=k1_xboole_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t56_relat_1])])])])).
cnf(c_0_19, plain, (r2_hidden(esk1_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_20, plain, ![X7, X8]:((~r1_xboole_0(X7,X8)|k4_xboole_0(X7,X8)=X7)&(k4_xboole_0(X7,X8)!=X7|r1_xboole_0(X7,X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
fof(c_0_21, plain, ![X9, X10]:((k4_xboole_0(X9,X10)!=k1_xboole_0|r1_tarski(X9,X10))&(~r1_tarski(X9,X10)|k4_xboole_0(X9,X10)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
fof(c_0_22, plain, ![X17, X18]:(((r1_tarski(X17,X18)|X17!=X18)&(r1_tarski(X18,X17)|X17!=X18))&(~r1_tarski(X17,X18)|~r1_tarski(X18,X17)|X17=X18)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_23, plain, (~v1_relat_1(X1)|~v1_relat_1(X2)|~r2_hidden(k4_tarski(X3,esk2_5(X2,X1,k5_relat_1(X2,X1),X3,X4)),X5)|~r2_hidden(k4_tarski(X3,X4),k5_relat_1(X2,X1))|~r1_xboole_0(X5,X2)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_24, plain, (r2_hidden(k4_tarski(esk6_1(X1),esk7_1(X1)),X1)|X1=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (r1_xboole_0(X1,X2)|~r2_hidden(esk1_2(X1,X2),X3)|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_16, c_0_19])).
cnf(c_0_26, plain, (r1_xboole_0(X1,X2)|k4_xboole_0(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_29, plain, ![X37, X38]:(~v1_relat_1(X37)|v1_relat_1(k3_xboole_0(X37,X38))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_relat_1])])).
fof(c_0_30, plain, ![X39, X40]:k1_setfam_1(k2_tarski(X39,X40))=k3_xboole_0(X39,X40), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_31, plain, ![X41]:k3_xboole_0(X41,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
cnf(c_0_32, plain, (~v1_relat_1(X1)|~v1_relat_1(X2)|~r2_hidden(k4_tarski(X3,X4),k5_relat_1(X2,X1))|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_23, c_0_17])).
cnf(c_0_33, plain, (k5_relat_1(X1,X2)=k1_xboole_0|r2_hidden(k4_tarski(esk6_1(k5_relat_1(X1,X2)),esk7_1(k5_relat_1(X1,X2))),k5_relat_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_24, c_0_15])).
cnf(c_0_34, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_25, c_0_19])).
cnf(c_0_35, plain, (r1_xboole_0(X1,X2)|k1_xboole_0!=X1|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_36, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_28])).
fof(c_0_37, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>(k5_relat_1(k1_xboole_0,X1)=k1_xboole_0&k5_relat_1(X1,k1_xboole_0)=k1_xboole_0))), inference(assume_negation,[status(cth)],[t62_relat_1])).
cnf(c_0_38, plain, (v1_relat_1(k3_xboole_0(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41, plain, (k5_relat_1(X1,X2)=k1_xboole_0|~v1_relat_1(X2)|~v1_relat_1(X1)|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_42, plain, (r1_xboole_0(X1,X2)|k1_xboole_0!=X2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])])).
fof(c_0_43, negated_conjecture, (v1_relat_1(esk8_0)&(k5_relat_1(k1_xboole_0,esk8_0)!=k1_xboole_0|k5_relat_1(esk8_0,k1_xboole_0)!=k1_xboole_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_37])])])).
cnf(c_0_44, plain, (v1_relat_1(k1_setfam_1(k2_tarski(X1,X2)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_45, plain, (k1_setfam_1(k2_tarski(X1,k1_xboole_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_40, c_0_39])).
cnf(c_0_46, plain, (k5_relat_1(X1,X2)=k1_xboole_0|k1_xboole_0!=X1|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_47, negated_conjecture, (v1_relat_1(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_48, plain, (v1_relat_1(k1_xboole_0)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_49, negated_conjecture, (k5_relat_1(X1,esk8_0)=k1_xboole_0|k1_xboole_0!=X1|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_50, negated_conjecture, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_48, c_0_47])).
cnf(c_0_51, negated_conjecture, (k5_relat_1(k1_xboole_0,esk8_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_52, negated_conjecture, (~r2_hidden(k4_tarski(X1,X2),k1_xboole_0)|~r1_xboole_0(k1_xboole_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_51]), c_0_47]), c_0_50])])).
cnf(c_0_53, plain, (r2_hidden(k4_tarski(esk2_5(X1,X2,X3,X4,X5),X5),X2)|~r2_hidden(k4_tarski(X4,X5),X3)|X3!=k5_relat_1(X1,X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_54, negated_conjecture, (~r2_hidden(k4_tarski(X1,X2),k1_xboole_0)), inference(spm,[status(thm)],[c_0_52, c_0_42])).
cnf(c_0_55, plain, (r2_hidden(k4_tarski(esk2_5(X1,X2,k5_relat_1(X1,X2),X3,X4),X4),X2)|~v1_relat_1(X2)|~v1_relat_1(X1)|~r2_hidden(k4_tarski(X3,X4),k5_relat_1(X1,X2))), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_53]), c_0_15])).
cnf(c_0_56, negated_conjecture, (~v1_relat_1(X1)|~r2_hidden(k4_tarski(X2,X3),k5_relat_1(X1,k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_50])])).
cnf(c_0_57, negated_conjecture, (k5_relat_1(k1_xboole_0,esk8_0)!=k1_xboole_0|k5_relat_1(esk8_0,k1_xboole_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_58, negated_conjecture, (k5_relat_1(X1,k1_xboole_0)=k1_xboole_0|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_33]), c_0_50])])).
cnf(c_0_59, negated_conjecture, (k5_relat_1(esk8_0,k1_xboole_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_51])])).
cnf(c_0_60, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_47]), c_0_59]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 61
# Proof object clause steps            : 38
# Proof object formula steps           : 23
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 22
# Processed clauses                    : 113
# ...of these trivial                  : 2
# ...subsumed                          : 15
# ...remaining for further processing  : 96
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 2
# Generated clauses                    : 184
# ...of the previous two non-trivial   : 148
# Contextual simplify-reflections      : 2
# Paramodulations                      : 179
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 66
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 55
# Current number of unprocessed clauses: 74
# ...number of literals in the above   : 392
# Current number of archived formulas  : 0
# Current number of archived clauses   : 29
# Clause-clause subsumption calls (NU) : 1717
# Rec. Clause-clause subsumption calls : 511
# Non-unit clause-clause subsumptions  : 17
# Unit Clause-clause subsumption calls : 47
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5765

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28720 pages
