# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/4b2455e14bc4486dbd4a98a0c7942487.E_input', t12_setfam_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/4b2455e14bc4486dbd4a98a0c7942487.E_input', t70_enumset1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/4b2455e14bc4486dbd4a98a0c7942487.E_input', d1_tarski)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/4b2455e14bc4486dbd4a98a0c7942487.E_input', t69_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/4b2455e14bc4486dbd4a98a0c7942487.E_input', t71_enumset1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/4b2455e14bc4486dbd4a98a0c7942487.E_input', t100_xboole_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/4b2455e14bc4486dbd4a98a0c7942487.E_input', d5_xboole_0)).
fof(t20_mcart_1, conjecture, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1)))), file('data/MPTP/proofs/4b2455e14bc4486dbd4a98a0c7942487.E_input', t20_mcart_1)).
fof(t9_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k4_tarski(X3,X4))))))), file('data/MPTP/proofs/4b2455e14bc4486dbd4a98a0c7942487.E_input', t9_mcart_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/4b2455e14bc4486dbd4a98a0c7942487.E_input', idempotence_k3_xboole_0)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/4b2455e14bc4486dbd4a98a0c7942487.E_input', t1_boole)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/4b2455e14bc4486dbd4a98a0c7942487.E_input', t98_xboole_1)).
fof(t4_boole, axiom, ![X1]:k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0, file('data/MPTP/proofs/4b2455e14bc4486dbd4a98a0c7942487.E_input', t4_boole)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/4b2455e14bc4486dbd4a98a0c7942487.E_input', t7_mcart_1)).
fof(d1_enumset1, axiom, ![X1, X2, X3, X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3)))))), file('data/MPTP/proofs/4b2455e14bc4486dbd4a98a0c7942487.E_input', d1_enumset1)).
fof(c_0_15, plain, ![X41, X42]:k1_setfam_1(k2_tarski(X41,X42))=k3_xboole_0(X41,X42), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_16, plain, ![X36, X37]:k1_enumset1(X36,X36,X37)=k2_tarski(X36,X37), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_17, plain, ![X27, X28, X29, X30, X31, X32]:(((~r2_hidden(X29,X28)|X29=X27|X28!=k1_tarski(X27))&(X30!=X27|r2_hidden(X30,X28)|X28!=k1_tarski(X27)))&((~r2_hidden(esk6_2(X31,X32),X32)|esk6_2(X31,X32)!=X31|X32=k1_tarski(X31))&(r2_hidden(esk6_2(X31,X32),X32)|esk6_2(X31,X32)=X31|X32=k1_tarski(X31)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_18, plain, ![X24]:k2_tarski(X24,X24)=k1_tarski(X24), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_19, plain, ![X38, X39, X40]:k2_enumset1(X38,X38,X39,X40)=k1_enumset1(X38,X39,X40), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_20, plain, ![X44, X45]:k4_xboole_0(X44,X45)=k5_xboole_0(X44,k3_xboole_0(X44,X45)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_21, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_26, plain, ![X48, X49, X50, X51, X52, X53, X54, X55]:((((r2_hidden(X51,X48)|~r2_hidden(X51,X50)|X50!=k4_xboole_0(X48,X49))&(~r2_hidden(X51,X49)|~r2_hidden(X51,X50)|X50!=k4_xboole_0(X48,X49)))&(~r2_hidden(X52,X48)|r2_hidden(X52,X49)|r2_hidden(X52,X50)|X50!=k4_xboole_0(X48,X49)))&((~r2_hidden(esk7_3(X53,X54,X55),X55)|(~r2_hidden(esk7_3(X53,X54,X55),X53)|r2_hidden(esk7_3(X53,X54,X55),X54))|X55=k4_xboole_0(X53,X54))&((r2_hidden(esk7_3(X53,X54,X55),X53)|r2_hidden(esk7_3(X53,X54,X55),X55)|X55=k4_xboole_0(X53,X54))&(~r2_hidden(esk7_3(X53,X54,X55),X54)|r2_hidden(esk7_3(X53,X54,X55),X55)|X55=k4_xboole_0(X53,X54))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_27, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_21, c_0_22])).
fof(c_0_29, negated_conjecture, ~(![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1))))), inference(assume_negation,[status(cth)],[t20_mcart_1])).
fof(c_0_30, plain, ![X9, X11, X12]:((r2_hidden(esk4_1(X9),X9)|X9=k1_xboole_0)&((~r2_hidden(X11,X9)|esk4_1(X9)!=k4_tarski(X11,X12)|X9=k1_xboole_0)&(~r2_hidden(X12,X9)|esk4_1(X9)!=k4_tarski(X11,X12)|X9=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_mcart_1])])])])])])).
cnf(c_0_31, plain, (X1=X3|X2!=k2_enumset1(X3,X3,X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_24]), c_0_22]), c_0_25])).
cnf(c_0_32, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[c_0_27, c_0_28])).
fof(c_0_34, plain, ![X43]:k3_xboole_0(X43,X43)=X43, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
fof(c_0_35, plain, ![X47]:k2_xboole_0(X47,k1_xboole_0)=X47, inference(variable_rename,[status(thm)],[t1_boole])).
fof(c_0_36, plain, ![X25, X26]:k2_xboole_0(X25,X26)=k5_xboole_0(X25,k4_xboole_0(X26,X25)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_37, plain, ![X46]:k4_xboole_0(k1_xboole_0,X46)=k1_xboole_0, inference(variable_rename,[status(thm)],[t4_boole])).
fof(c_0_38, plain, ![X34, X35]:(k1_mcart_1(k4_tarski(X34,X35))=X34&k2_mcart_1(k4_tarski(X34,X35))=X35), inference(variable_rename,[status(thm)],[t7_mcart_1])).
fof(c_0_39, negated_conjecture, (esk1_0=k4_tarski(esk2_0,esk3_0)&(esk1_0=k1_mcart_1(esk1_0)|esk1_0=k2_mcart_1(esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_29])])])])).
fof(c_0_40, plain, ![X13, X14, X15, X16, X17, X18, X19, X20, X21, X22]:(((~r2_hidden(X17,X16)|(X17=X13|(X17=X14|X17=X15))|X16!=k1_enumset1(X13,X14,X15))&((X18!=X13|r2_hidden(X18,X16)|X16!=k1_enumset1(X13,X14,X15))&((X18!=X14|r2_hidden(X18,X16)|X16!=k1_enumset1(X13,X14,X15))&(X18!=X15|r2_hidden(X18,X16)|X16!=k1_enumset1(X13,X14,X15)))))&(((esk5_4(X19,X20,X21,X22)!=X19|~r2_hidden(esk5_4(X19,X20,X21,X22),X22)|X22=k1_enumset1(X19,X20,X21))&((esk5_4(X19,X20,X21,X22)!=X20|~r2_hidden(esk5_4(X19,X20,X21,X22),X22)|X22=k1_enumset1(X19,X20,X21))&(esk5_4(X19,X20,X21,X22)!=X21|~r2_hidden(esk5_4(X19,X20,X21,X22),X22)|X22=k1_enumset1(X19,X20,X21))))&(r2_hidden(esk5_4(X19,X20,X21,X22),X22)|(esk5_4(X19,X20,X21,X22)=X19|(esk5_4(X19,X20,X21,X22)=X20|esk5_4(X19,X20,X21,X22)=X21))|X22=k1_enumset1(X19,X20,X21)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])).
cnf(c_0_41, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk4_1(X2)!=k4_tarski(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_42, plain, (r2_hidden(esk4_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_43, plain, (X1=X2|~r2_hidden(X1,k2_enumset1(X2,X2,X2,X2))), inference(er,[status(thm)],[c_0_31])).
cnf(c_0_44, plain, (X3!=k5_xboole_0(X4,k1_setfam_1(k2_enumset1(X4,X4,X4,X2)))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_33]), c_0_25])).
cnf(c_0_45, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_46, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_47, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_48, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_49, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, negated_conjecture, (esk1_0=k4_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X2,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_52, plain, (X1=k1_xboole_0|k4_tarski(X2,esk4_1(X1))!=esk4_1(X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_53, plain, (esk4_1(k2_enumset1(X1,X1,X1,X1))=X1|k2_enumset1(X1,X1,X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_43, c_0_42])).
cnf(c_0_54, plain, (~r2_hidden(X1,k5_xboole_0(X2,k1_setfam_1(k2_enumset1(X2,X2,X2,X3))))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_55, plain, (k1_setfam_1(k2_enumset1(X1,X1,X1,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_28]), c_0_25])).
cnf(c_0_56, plain, (k5_xboole_0(X1,k5_xboole_0(k1_xboole_0,k1_setfam_1(k2_enumset1(k1_xboole_0,k1_xboole_0,k1_xboole_0,X1))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_47]), c_0_33]), c_0_25])).
cnf(c_0_57, plain, (k5_xboole_0(k1_xboole_0,k1_setfam_1(k2_enumset1(k1_xboole_0,k1_xboole_0,k1_xboole_0,X1)))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_33]), c_0_25])).
cnf(c_0_58, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_59, negated_conjecture, (esk1_0=k1_mcart_1(esk1_0)|esk1_0=k2_mcart_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_60, negated_conjecture, (k1_mcart_1(esk1_0)=esk2_0), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_61, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_enumset1(X2,X2,X4,X5)), inference(rw,[status(thm)],[c_0_51, c_0_25])).
cnf(c_0_62, plain, (k2_enumset1(X1,X1,X1,X1)=k1_xboole_0|k4_tarski(X2,X1)!=X1), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_63, plain, (~r2_hidden(X1,k5_xboole_0(X2,X2))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_64, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_65, negated_conjecture, (k2_mcart_1(esk1_0)=esk3_0), inference(spm,[status(thm)],[c_0_58, c_0_50])).
cnf(c_0_66, negated_conjecture, (k2_mcart_1(esk1_0)=esk1_0|esk2_0=esk1_0), inference(rw,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_67, plain, (r2_hidden(X1,k2_enumset1(X1,X1,X2,X3))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_61])])).
cnf(c_0_68, negated_conjecture, (k2_enumset1(esk3_0,esk3_0,esk3_0,esk3_0)=k1_xboole_0|esk3_0!=esk1_0), inference(spm,[status(thm)],[c_0_62, c_0_50])).
cnf(c_0_69, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_70, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk4_1(X2)!=k4_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_71, negated_conjecture, (esk2_0=esk1_0|esk3_0=esk1_0), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_72, negated_conjecture, (esk3_0!=esk1_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_69])).
cnf(c_0_73, plain, (X1=k1_xboole_0|k4_tarski(esk4_1(X1),X2)!=esk4_1(X1)), inference(spm,[status(thm)],[c_0_70, c_0_42])).
cnf(c_0_74, negated_conjecture, (esk2_0=esk1_0), inference(sr,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_75, plain, (k2_enumset1(X1,X1,X1,X1)=k1_xboole_0|k4_tarski(X1,X2)!=X1), inference(spm,[status(thm)],[c_0_73, c_0_53])).
cnf(c_0_76, negated_conjecture, (k4_tarski(esk1_0,esk3_0)=esk1_0), inference(rw,[status(thm)],[c_0_50, c_0_74])).
cnf(c_0_77, negated_conjecture, (k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_78, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_77]), c_0_69]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 79
# Proof object clause steps            : 48
# Proof object formula steps           : 31
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 25
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 6
# Initial clauses in saturation        : 28
# Processed clauses                    : 117
# ...of these trivial                  : 1
# ...subsumed                          : 17
# ...remaining for further processing  : 99
# Other redundant clauses eliminated   : 25
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 8
# Generated clauses                    : 373
# ...of the previous two non-trivial   : 337
# Contextual simplify-reflections      : 1
# Paramodulations                      : 351
# Factorizations                       : 0
# Equation resolutions                 : 25
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 53
#    Positive orientable unit clauses  : 15
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 36
# Current number of unprocessed clauses: 257
# ...number of literals in the above   : 1050
# Current number of archived formulas  : 0
# Current number of archived clauses   : 43
# Clause-clause subsumption calls (NU) : 186
# Rec. Clause-clause subsumption calls : 125
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 42
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 21
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7015

# -------------------------------------------------
# User time                : 0.014 s
# System time              : 0.000 s
# Total time               : 0.014 s
# Maximum resident set size: 28596 pages
