# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(free_g1_orders_2, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=>![X3, X4]:(g1_orders_2(X1,X2)=g1_orders_2(X3,X4)=>(X1=X3&X2=X4))), file('data/MPTP/proofs/4b66bdb27f9a4e3d8e0a504bf3fd8989.E_input', free_g1_orders_2)).
fof(dt_u1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))), file('data/MPTP/proofs/4b66bdb27f9a4e3d8e0a504bf3fd8989.E_input', dt_u1_orders_2)).
fof(t26_yellow_0, conjecture, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:(r1_yellow_0(X1,X3)=>k1_yellow_0(X1,X3)=k1_yellow_0(X2,X3))))), file('data/MPTP/proofs/4b66bdb27f9a4e3d8e0a504bf3fd8989.E_input', t26_yellow_0)).
fof(t1_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>((X3=X5&X4=X6)=>((r1_orders_2(X1,X3,X4)=>r1_orders_2(X2,X5,X6))&(r2_orders_2(X1,X3,X4)=>r2_orders_2(X2,X5,X6))))))))))), file('data/MPTP/proofs/4b66bdb27f9a4e3d8e0a504bf3fd8989.E_input', t1_yellow_0)).
fof(d7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X5)=>r1_orders_2(X1,X4,X5))))=>X4=X3))))))), file('data/MPTP/proofs/4b66bdb27f9a4e3d8e0a504bf3fd8989.E_input', d7_yellow_0)).
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/4b66bdb27f9a4e3d8e0a504bf3fd8989.E_input', d9_yellow_0)).
fof(t2_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3, X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(X4=X5=>((r2_lattice3(X1,X3,X4)=>r2_lattice3(X2,X3,X5))&(r1_lattice3(X1,X3,X4)=>r1_lattice3(X2,X3,X5))))))))), file('data/MPTP/proofs/4b66bdb27f9a4e3d8e0a504bf3fd8989.E_input', t2_yellow_0)).
fof(c_0_7, plain, ![X7, X8, X9, X10]:((X7=X9|g1_orders_2(X7,X8)!=g1_orders_2(X9,X10)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X7,X7))))&(X8=X10|g1_orders_2(X7,X8)!=g1_orders_2(X9,X10)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X7,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_orders_2])])])])).
fof(c_0_8, plain, ![X11]:(~l1_orders_2(X11)|m1_subset_1(u1_orders_2(X11),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X11),u1_struct_0(X11))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_orders_2])])).
fof(c_0_9, negated_conjecture, ~(![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:(r1_yellow_0(X1,X3)=>k1_yellow_0(X1,X3)=k1_yellow_0(X2,X3)))))), inference(assume_negation,[status(cth)],[t26_yellow_0])).
cnf(c_0_10, plain, (X1=X2|g1_orders_2(X3,X1)!=g1_orders_2(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X3)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_12, negated_conjecture, (l1_orders_2(esk1_0)&(l1_orders_2(esk2_0)&(g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))&(r1_yellow_0(esk1_0,esk3_0)&k1_yellow_0(esk1_0,esk3_0)!=k1_yellow_0(esk2_0,esk3_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
cnf(c_0_13, plain, (u1_orders_2(X1)=X2|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(X3,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_14, negated_conjecture, (g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15, negated_conjecture, (l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, negated_conjecture, (u1_orders_2(esk2_0)=X1|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])])).
fof(c_0_17, plain, ![X15, X16, X17, X18, X19, X20]:((~r1_orders_2(X15,X17,X18)|r1_orders_2(X16,X19,X20)|(X17!=X19|X18!=X20)|~m1_subset_1(X20,u1_struct_0(X16))|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X15))|~m1_subset_1(X17,u1_struct_0(X15))|g1_orders_2(u1_struct_0(X15),u1_orders_2(X15))!=g1_orders_2(u1_struct_0(X16),u1_orders_2(X16))|~l1_orders_2(X16)|~l1_orders_2(X15))&(~r2_orders_2(X15,X17,X18)|r2_orders_2(X16,X19,X20)|(X17!=X19|X18!=X20)|~m1_subset_1(X20,u1_struct_0(X16))|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X15))|~m1_subset_1(X17,u1_struct_0(X15))|g1_orders_2(u1_struct_0(X15),u1_orders_2(X15))!=g1_orders_2(u1_struct_0(X16),u1_orders_2(X16))|~l1_orders_2(X16)|~l1_orders_2(X15))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_yellow_0])])])])).
cnf(c_0_18, negated_conjecture, (u1_orders_2(esk2_0)=u1_orders_2(esk1_0)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_19, plain, (r1_orders_2(X4,X5,X6)|~r1_orders_2(X1,X2,X3)|X2!=X5|X3!=X6|~m1_subset_1(X6,u1_struct_0(X4))|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))|~l1_orders_2(X4)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_20, plain, ![X21, X22, X24, X25, X27, X28, X31]:(((m1_subset_1(esk4_2(X21,X22),u1_struct_0(X21))|~r1_yellow_0(X21,X22)|~l1_orders_2(X21))&((r2_lattice3(X21,X22,esk4_2(X21,X22))|~r1_yellow_0(X21,X22)|~l1_orders_2(X21))&((~m1_subset_1(X24,u1_struct_0(X21))|(~r2_lattice3(X21,X22,X24)|r1_orders_2(X21,esk4_2(X21,X22),X24))|~r1_yellow_0(X21,X22)|~l1_orders_2(X21))&((m1_subset_1(esk5_3(X21,X22,X25),u1_struct_0(X21))|~r2_lattice3(X21,X22,X25)|X25=esk4_2(X21,X22)|~m1_subset_1(X25,u1_struct_0(X21))|~r1_yellow_0(X21,X22)|~l1_orders_2(X21))&((r2_lattice3(X21,X22,esk5_3(X21,X22,X25))|~r2_lattice3(X21,X22,X25)|X25=esk4_2(X21,X22)|~m1_subset_1(X25,u1_struct_0(X21))|~r1_yellow_0(X21,X22)|~l1_orders_2(X21))&(~r1_orders_2(X21,X25,esk5_3(X21,X22,X25))|~r2_lattice3(X21,X22,X25)|X25=esk4_2(X21,X22)|~m1_subset_1(X25,u1_struct_0(X21))|~r1_yellow_0(X21,X22)|~l1_orders_2(X21)))))))&(((m1_subset_1(esk7_3(X21,X27,X28),u1_struct_0(X21))|m1_subset_1(esk6_3(X21,X27,X28),u1_struct_0(X21))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21))&(((r2_lattice3(X21,X27,esk7_3(X21,X27,X28))|m1_subset_1(esk6_3(X21,X27,X28),u1_struct_0(X21))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21))&(~m1_subset_1(X31,u1_struct_0(X21))|(~r2_lattice3(X21,X27,X31)|r1_orders_2(X21,esk7_3(X21,X27,X28),X31))|m1_subset_1(esk6_3(X21,X27,X28),u1_struct_0(X21))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21)))&(esk7_3(X21,X27,X28)!=X28|m1_subset_1(esk6_3(X21,X27,X28),u1_struct_0(X21))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21))))&(((m1_subset_1(esk7_3(X21,X27,X28),u1_struct_0(X21))|r2_lattice3(X21,X27,esk6_3(X21,X27,X28))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21))&(((r2_lattice3(X21,X27,esk7_3(X21,X27,X28))|r2_lattice3(X21,X27,esk6_3(X21,X27,X28))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21))&(~m1_subset_1(X31,u1_struct_0(X21))|(~r2_lattice3(X21,X27,X31)|r1_orders_2(X21,esk7_3(X21,X27,X28),X31))|r2_lattice3(X21,X27,esk6_3(X21,X27,X28))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21)))&(esk7_3(X21,X27,X28)!=X28|r2_lattice3(X21,X27,esk6_3(X21,X27,X28))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21))))&((m1_subset_1(esk7_3(X21,X27,X28),u1_struct_0(X21))|~r1_orders_2(X21,X28,esk6_3(X21,X27,X28))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21))&(((r2_lattice3(X21,X27,esk7_3(X21,X27,X28))|~r1_orders_2(X21,X28,esk6_3(X21,X27,X28))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21))&(~m1_subset_1(X31,u1_struct_0(X21))|(~r2_lattice3(X21,X27,X31)|r1_orders_2(X21,esk7_3(X21,X27,X28),X31))|~r1_orders_2(X21,X28,esk6_3(X21,X27,X28))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21)))&(esk7_3(X21,X27,X28)!=X28|~r1_orders_2(X21,X28,esk6_3(X21,X27,X28))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_yellow_0])])])])])])).
cnf(c_0_21, plain, (X1=X2|g1_orders_2(X1,X3)!=g1_orders_2(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_22, negated_conjecture, (m1_subset_1(u1_orders_2(esk1_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_18]), c_0_15])])).
cnf(c_0_23, negated_conjecture, (g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))), inference(rw,[status(thm)],[c_0_14, c_0_18])).
fof(c_0_24, plain, ![X32, X33, X34, X35]:(((r2_lattice3(X32,X33,X34)|X34!=k1_yellow_0(X32,X33)|~r1_yellow_0(X32,X33)|~m1_subset_1(X34,u1_struct_0(X32))|~l1_orders_2(X32))&(~m1_subset_1(X35,u1_struct_0(X32))|(~r2_lattice3(X32,X33,X35)|r1_orders_2(X32,X34,X35))|X34!=k1_yellow_0(X32,X33)|~r1_yellow_0(X32,X33)|~m1_subset_1(X34,u1_struct_0(X32))|~l1_orders_2(X32)))&((m1_subset_1(esk8_3(X32,X33,X34),u1_struct_0(X32))|~r2_lattice3(X32,X33,X34)|X34=k1_yellow_0(X32,X33)|~r1_yellow_0(X32,X33)|~m1_subset_1(X34,u1_struct_0(X32))|~l1_orders_2(X32))&((r2_lattice3(X32,X33,esk8_3(X32,X33,X34))|~r2_lattice3(X32,X33,X34)|X34=k1_yellow_0(X32,X33)|~r1_yellow_0(X32,X33)|~m1_subset_1(X34,u1_struct_0(X32))|~l1_orders_2(X32))&(~r1_orders_2(X32,X34,esk8_3(X32,X33,X34))|~r2_lattice3(X32,X33,X34)|X34=k1_yellow_0(X32,X33)|~r1_yellow_0(X32,X33)|~m1_subset_1(X34,u1_struct_0(X32))|~l1_orders_2(X32))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
fof(c_0_25, plain, ![X37, X38, X39, X40, X41]:((~r2_lattice3(X37,X39,X40)|r2_lattice3(X38,X39,X41)|X40!=X41|~m1_subset_1(X41,u1_struct_0(X38))|~m1_subset_1(X40,u1_struct_0(X37))|g1_orders_2(u1_struct_0(X37),u1_orders_2(X37))!=g1_orders_2(u1_struct_0(X38),u1_orders_2(X38))|~l1_orders_2(X38)|~l1_orders_2(X37))&(~r1_lattice3(X37,X39,X40)|r1_lattice3(X38,X39,X41)|X40!=X41|~m1_subset_1(X41,u1_struct_0(X38))|~m1_subset_1(X40,u1_struct_0(X37))|g1_orders_2(u1_struct_0(X37),u1_orders_2(X37))!=g1_orders_2(u1_struct_0(X38),u1_orders_2(X38))|~l1_orders_2(X38)|~l1_orders_2(X37))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_yellow_0])])])])).
cnf(c_0_26, plain, (r1_orders_2(X1,X2,X3)|g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_orders_2(X4,X2,X3)|~l1_orders_2(X1)|~l1_orders_2(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_19])])).
cnf(c_0_27, plain, (r1_orders_2(X2,esk4_2(X2,X3),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r1_yellow_0(X2,X3)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (m1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, negated_conjecture, (u1_struct_0(esk2_0)=X1|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])).
cnf(c_0_30, plain, (X2=k1_yellow_0(X1,X3)|~r1_orders_2(X1,X2,esk8_3(X1,X3,X2))|~r2_lattice3(X1,X3,X2)|~r1_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, plain, (r2_lattice3(X4,X2,X5)|~r2_lattice3(X1,X2,X3)|X3!=X5|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))|~l1_orders_2(X4)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, plain, (r1_orders_2(X1,esk4_2(X2,X3),X4)|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(X2,X3,X4)|~r1_yellow_0(X2,X3)|~l1_orders_2(X1)|~l1_orders_2(X2)|~m1_subset_1(esk4_2(X2,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])).
cnf(c_0_33, negated_conjecture, (u1_struct_0(esk2_0)=u1_struct_0(esk1_0)), inference(er,[status(thm)],[c_0_29])).
cnf(c_0_34, plain, (esk4_2(X1,X2)=k1_yellow_0(X1,X3)|~r2_lattice3(X1,X2,esk8_3(X1,X3,esk4_2(X1,X2)))|~r2_lattice3(X1,X3,esk4_2(X1,X2))|~r1_yellow_0(X1,X3)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)|~m1_subset_1(esk8_3(X1,X3,esk4_2(X1,X2)),u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_27]), c_0_28])).
cnf(c_0_35, plain, (r2_lattice3(X1,X2,esk8_3(X1,X2,X3))|X3=k1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_36, plain, (r2_lattice3(X1,X2,esk4_2(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_37, plain, (r2_lattice3(X1,X2,X3)|g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(X4,X2,X3)|~l1_orders_2(X1)|~l1_orders_2(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X4))), inference(er,[status(thm)],[c_0_31])).
cnf(c_0_38, negated_conjecture, (r1_orders_2(esk2_0,esk4_2(X1,X2),X3)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)|~m1_subset_1(esk4_2(X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_18]), c_0_33]), c_0_15]), c_0_33]), c_0_33])])).
cnf(c_0_39, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_40, plain, (esk4_2(X1,X2)=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)|~m1_subset_1(esk8_3(X1,X2,esk4_2(X1,X2)),u1_struct_0(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_28]), c_0_36])).
cnf(c_0_41, plain, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|X3=k1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_42, plain, (r2_lattice3(X1,X2,esk4_2(X3,X2))|g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_yellow_0(X3,X2)|~l1_orders_2(X1)|~l1_orders_2(X3)|~m1_subset_1(esk4_2(X3,X2),u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_36]), c_0_28])).
cnf(c_0_43, negated_conjecture, (r1_orders_2(esk2_0,esk4_2(esk1_0,X1),X2)|~r2_lattice3(esk1_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk4_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_38]), c_0_39])])).
cnf(c_0_44, plain, (esk4_2(X1,X2)=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_28]), c_0_36])).
cnf(c_0_45, negated_conjecture, (r2_lattice3(esk2_0,X1,esk4_2(X2,X1))|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~r1_yellow_0(X2,X1)|~l1_orders_2(X2)|~m1_subset_1(esk4_2(X2,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_18]), c_0_33]), c_0_15]), c_0_33])])).
cnf(c_0_46, negated_conjecture, (r1_orders_2(esk2_0,k1_yellow_0(esk1_0,X1),X2)|~r2_lattice3(esk1_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_39])])).
cnf(c_0_47, negated_conjecture, (r2_lattice3(esk2_0,X1,esk4_2(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk4_2(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_45]), c_0_39])])).
cnf(c_0_48, negated_conjecture, (r1_orders_2(X1,k1_yellow_0(esk1_0,X2),X3)|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(esk1_0,X2,X3)|~r1_yellow_0(esk1_0,X2)|~l1_orders_2(X1)|~m1_subset_1(k1_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_46]), c_0_33]), c_0_18]), c_0_15]), c_0_33]), c_0_33])])).
cnf(c_0_49, negated_conjecture, (r2_lattice3(esk2_0,X1,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_44]), c_0_39])])).
cnf(c_0_50, plain, (r1_orders_2(X2,esk7_3(X2,X3,X4),X1)|r2_lattice3(X2,X3,esk6_3(X2,X3,X4))|r1_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r2_lattice3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_51, plain, (X2=esk4_2(X1,X3)|~r1_orders_2(X1,X2,esk5_3(X1,X3,X2))|~r2_lattice3(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_yellow_0(X1,X3)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_52, negated_conjecture, (r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),X2)|~r2_lattice3(esk1_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_48]), c_0_39])])).
cnf(c_0_53, negated_conjecture, (r2_lattice3(X1,X2,k1_yellow_0(esk1_0,X2))|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_yellow_0(esk1_0,X2)|~l1_orders_2(X1)|~m1_subset_1(k1_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,X2),u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_49]), c_0_33]), c_0_18]), c_0_15]), c_0_33])])).
cnf(c_0_54, negated_conjecture, (r2_lattice3(esk2_0,X1,esk6_3(esk2_0,X1,X2))|r1_orders_2(esk2_0,esk7_3(esk2_0,X1,X2),X3)|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,X2)|~r2_lattice3(esk2_0,X1,X3)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_33]), c_0_15])])).
cnf(c_0_55, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_28, c_0_44])).
cnf(c_0_56, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r2_lattice3(X1,X2,esk6_3(X1,X2,X3))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_57, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_58, negated_conjecture, (k1_yellow_0(esk1_0,X1)=esk4_2(esk1_0,X2)|~r2_lattice3(esk1_0,X1,esk5_3(esk1_0,X2,k1_yellow_0(esk1_0,X1)))|~r2_lattice3(esk1_0,X2,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk5_3(esk1_0,X2,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_39])])).
cnf(c_0_59, plain, (r2_lattice3(X1,X2,esk5_3(X1,X2,X3))|X3=esk4_2(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_60, negated_conjecture, (r2_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_53]), c_0_39])])).
cnf(c_0_61, negated_conjecture, (r2_lattice3(esk2_0,X1,esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X2)))|r1_orders_2(esk2_0,esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X2)),X3)|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,k1_yellow_0(esk1_0,X2))|~r2_lattice3(esk2_0,X1,X3)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(X3,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_39])])).
cnf(c_0_62, plain, (r2_lattice3(X1,X2,esk6_3(X3,X2,X4))|r1_yellow_0(X3,X2)|m1_subset_1(esk7_3(X3,X2,X4),u1_struct_0(X3))|g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(X3,X2,X4)|~l1_orders_2(X1)|~l1_orders_2(X3)|~m1_subset_1(esk6_3(X3,X2,X4),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X3))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_56]), c_0_57])).
cnf(c_0_63, negated_conjecture, (esk4_2(esk1_0,X1)=k1_yellow_0(esk1_0,X1)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk5_3(esk1_0,X1,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_39])]), c_0_60])).
cnf(c_0_64, plain, (m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(X1))|X3=esk4_2(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_65, negated_conjecture, (r2_lattice3(esk2_0,X1,esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)))|r1_orders_2(esk2_0,esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)),X2)|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_61, c_0_49])).
cnf(c_0_66, plain, (X1=k1_yellow_0(X2,X3)|r2_lattice3(X4,X3,esk8_3(X2,X3,X1))|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))|~r2_lattice3(X2,X3,X1)|~r1_yellow_0(X2,X3)|~l1_orders_2(X4)|~l1_orders_2(X2)|~m1_subset_1(esk8_3(X2,X3,X1),u1_struct_0(X4))|~m1_subset_1(X1,u1_struct_0(X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_35]), c_0_41])).
cnf(c_0_67, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|~r1_orders_2(X1,X3,esk6_3(X1,X2,X3))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_68, negated_conjecture, (r1_yellow_0(esk2_0,X1)|m1_subset_1(esk7_3(esk2_0,X1,X2),u1_struct_0(esk1_0))|m1_subset_1(esk6_3(esk2_0,X1,X2),u1_struct_0(esk1_0))|~r2_lattice3(esk2_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_33]), c_0_15])])).
cnf(c_0_69, negated_conjecture, (r2_lattice3(X1,X2,esk6_3(esk2_0,X2,X3))|r1_yellow_0(esk2_0,X2)|m1_subset_1(esk7_3(esk2_0,X2,X3),u1_struct_0(esk1_0))|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(esk2_0,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(esk6_3(esk2_0,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_18]), c_0_33]), c_0_33]), c_0_15]), c_0_33])])).
cnf(c_0_70, negated_conjecture, (esk4_2(esk1_0,X1)=k1_yellow_0(esk1_0,X1)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_39])]), c_0_60])).
cnf(c_0_71, plain, (r2_lattice3(X1,X2,esk6_3(X1,X2,X3))|r1_yellow_0(X1,X2)|esk7_3(X1,X2,X3)!=X3|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_72, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|esk7_3(X1,X2,X3)!=X3|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_73, negated_conjecture, (r2_lattice3(esk2_0,X1,esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)))|r1_orders_2(esk2_0,esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)),X2)|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_55]), c_0_39])])).
cnf(c_0_74, negated_conjecture, (X1=k1_yellow_0(X2,X3)|r2_lattice3(esk2_0,X3,esk8_3(X2,X3,X1))|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~r2_lattice3(X2,X3,X1)|~r1_yellow_0(X2,X3)|~l1_orders_2(X2)|~m1_subset_1(esk8_3(X2,X3,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_18]), c_0_33]), c_0_15]), c_0_33])])).
cnf(c_0_75, negated_conjecture, (r1_yellow_0(esk2_0,X1)|m1_subset_1(esk7_3(esk2_0,X1,esk4_2(esk1_0,X2)),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,X2,esk6_3(esk2_0,X1,esk4_2(esk1_0,X2)))|~r2_lattice3(esk2_0,X1,esk4_2(esk1_0,X2))|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(esk4_2(esk1_0,X2),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_43]), c_0_33]), c_0_15]), c_0_33])]), c_0_68])).
cnf(c_0_76, negated_conjecture, (r2_lattice3(esk1_0,X1,esk6_3(esk2_0,X1,X2))|r1_yellow_0(esk2_0,X1)|m1_subset_1(esk7_3(esk2_0,X1,X2),u1_struct_0(esk1_0))|~r2_lattice3(esk2_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_69]), c_0_39])]), c_0_68])).
cnf(c_0_77, negated_conjecture, (esk4_2(esk1_0,X1)=k1_yellow_0(esk1_0,X1)|~r1_yellow_0(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_55]), c_0_39])])).
cnf(c_0_78, plain, (r2_lattice3(X1,X2,esk6_3(X3,X2,X4))|r1_yellow_0(X3,X2)|g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|esk7_3(X3,X2,X4)!=X4|~r2_lattice3(X3,X2,X4)|~l1_orders_2(X1)|~l1_orders_2(X3)|~m1_subset_1(esk6_3(X3,X2,X4),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X3))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_71]), c_0_72])).
cnf(c_0_79, negated_conjecture, (X1=k1_yellow_0(esk1_0,X2)|r2_lattice3(esk2_0,X3,esk6_3(esk2_0,X3,k1_yellow_0(esk1_0,X3)))|r1_orders_2(esk2_0,esk7_3(esk2_0,X3,k1_yellow_0(esk1_0,X3)),esk8_3(esk1_0,X2,X1))|r1_yellow_0(esk2_0,X3)|~r2_lattice3(esk2_0,X3,esk8_3(esk1_0,X2,X1))|~r2_lattice3(esk1_0,X2,X1)|~r1_yellow_0(esk1_0,X3)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_41]), c_0_39])])).
cnf(c_0_80, negated_conjecture, (X1=k1_yellow_0(esk1_0,X2)|r2_lattice3(esk2_0,X2,esk8_3(esk1_0,X2,X1))|~r2_lattice3(esk1_0,X2,X1)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(esk8_3(esk1_0,X2,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_74]), c_0_39])])).
cnf(c_0_81, negated_conjecture, (r1_yellow_0(esk2_0,X1)|m1_subset_1(esk7_3(esk2_0,X1,esk4_2(esk1_0,X1)),u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk4_2(esk1_0,X1),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_47])).
cnf(c_0_82, negated_conjecture, (m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_77]), c_0_39])])).
cnf(c_0_83, plain, (r1_yellow_0(X1,X2)|esk7_3(X1,X2,X3)!=X3|~r1_orders_2(X1,X3,esk6_3(X1,X2,X3))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_84, negated_conjecture, (r1_yellow_0(esk2_0,X1)|m1_subset_1(esk6_3(esk2_0,X1,X2),u1_struct_0(esk1_0))|esk7_3(esk2_0,X1,X2)!=X2|~r2_lattice3(esk2_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_33]), c_0_15])])).
cnf(c_0_85, negated_conjecture, (r2_lattice3(X1,X2,esk6_3(esk2_0,X2,X3))|r1_yellow_0(esk2_0,X2)|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|esk7_3(esk2_0,X2,X3)!=X3|~r2_lattice3(esk2_0,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(esk6_3(esk2_0,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_18]), c_0_33]), c_0_15]), c_0_33])])).
cnf(c_0_86, negated_conjecture, (X1=k1_yellow_0(esk1_0,X2)|r2_lattice3(esk2_0,X2,esk6_3(esk2_0,X2,k1_yellow_0(esk1_0,X2)))|r1_orders_2(esk2_0,esk7_3(esk2_0,X2,k1_yellow_0(esk1_0,X2)),esk8_3(esk1_0,X2,X1))|r1_yellow_0(esk2_0,X2)|~r2_lattice3(esk1_0,X2,X1)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(esk8_3(esk1_0,X2,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_87, negated_conjecture, (r1_yellow_0(esk2_0,X1)|m1_subset_1(esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_44]), c_0_39])]), c_0_82])).
cnf(c_0_88, negated_conjecture, (r1_yellow_0(esk2_0,X1)|esk7_3(esk2_0,X1,esk4_2(esk1_0,X2))!=esk4_2(esk1_0,X2)|~r2_lattice3(esk1_0,X2,esk6_3(esk2_0,X1,esk4_2(esk1_0,X2)))|~r2_lattice3(esk2_0,X1,esk4_2(esk1_0,X2))|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(esk4_2(esk1_0,X2),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_43]), c_0_15]), c_0_33])]), c_0_84])).
cnf(c_0_89, negated_conjecture, (r2_lattice3(esk1_0,X1,esk6_3(esk2_0,X1,X2))|r1_yellow_0(esk2_0,X1)|esk7_3(esk2_0,X1,X2)!=X2|~r2_lattice3(esk2_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_85]), c_0_39])]), c_0_84])).
cnf(c_0_90, plain, (r2_lattice3(X1,X2,esk7_3(X1,X2,X3))|r2_lattice3(X1,X2,esk6_3(X1,X2,X3))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_91, plain, (r2_lattice3(X1,X2,esk7_3(X1,X2,X3))|m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_92, negated_conjecture, (X1=k1_yellow_0(esk1_0,X2)|r2_lattice3(esk2_0,X2,esk6_3(esk2_0,X2,k1_yellow_0(esk1_0,X2)))|r1_orders_2(X3,esk7_3(esk2_0,X2,k1_yellow_0(esk1_0,X2)),esk8_3(esk1_0,X2,X1))|r1_yellow_0(esk2_0,X2)|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))|~r2_lattice3(esk1_0,X2,X1)|~r1_yellow_0(esk1_0,X2)|~l1_orders_2(X3)|~m1_subset_1(esk7_3(esk2_0,X2,k1_yellow_0(esk1_0,X2)),u1_struct_0(X3))|~m1_subset_1(esk8_3(esk1_0,X2,X1),u1_struct_0(esk1_0))|~m1_subset_1(esk8_3(esk1_0,X2,X1),u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_86]), c_0_33]), c_0_18]), c_0_15]), c_0_33]), c_0_33])]), c_0_87])).
cnf(c_0_93, negated_conjecture, (r1_yellow_0(esk2_0,X1)|esk7_3(esk2_0,X1,esk4_2(esk1_0,X1))!=esk4_2(esk1_0,X1)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk4_2(esk1_0,X1),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_47])).
cnf(c_0_94, plain, (r2_lattice3(X1,X2,esk6_3(X1,X2,X3))|r2_lattice3(X4,X2,esk7_3(X1,X2,X3))|r1_yellow_0(X1,X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))|~r2_lattice3(X1,X2,X3)|~l1_orders_2(X4)|~l1_orders_2(X1)|~m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_90]), c_0_56])).
cnf(c_0_95, plain, (r2_lattice3(X1,X2,esk7_3(X3,X2,X4))|r1_yellow_0(X3,X2)|m1_subset_1(esk6_3(X3,X2,X4),u1_struct_0(X3))|g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(X3,X2,X4)|~l1_orders_2(X1)|~l1_orders_2(X3)|~m1_subset_1(esk7_3(X3,X2,X4),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X3))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_91]), c_0_57])).
cnf(c_0_96, negated_conjecture, (X1=k1_yellow_0(esk1_0,X2)|r2_lattice3(esk2_0,X2,esk6_3(esk2_0,X2,k1_yellow_0(esk1_0,X2)))|r1_orders_2(esk1_0,esk7_3(esk2_0,X2,k1_yellow_0(esk1_0,X2)),esk8_3(esk1_0,X2,X1))|r1_yellow_0(esk2_0,X2)|~r2_lattice3(esk1_0,X2,X1)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(esk8_3(esk1_0,X2,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_92]), c_0_39])]), c_0_87])).
cnf(c_0_97, negated_conjecture, (r1_yellow_0(esk2_0,X1)|esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X1))!=k1_yellow_0(esk1_0,X1)|~r1_yellow_0(esk1_0,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_44]), c_0_39])]), c_0_82])).
cnf(c_0_98, negated_conjecture, (r2_lattice3(esk2_0,X1,esk6_3(esk2_0,X1,X2))|r2_lattice3(X3,X1,esk7_3(esk2_0,X1,X2))|r1_yellow_0(esk2_0,X1)|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))|~r2_lattice3(esk2_0,X1,X2)|~l1_orders_2(X3)|~m1_subset_1(esk7_3(esk2_0,X1,X2),u1_struct_0(X3))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_18]), c_0_33]), c_0_15]), c_0_33])])).
cnf(c_0_99, negated_conjecture, (r2_lattice3(X1,X2,esk7_3(esk2_0,X2,X3))|r1_yellow_0(esk2_0,X2)|m1_subset_1(esk6_3(esk2_0,X2,X3),u1_struct_0(esk1_0))|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(esk2_0,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(esk7_3(esk2_0,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_18]), c_0_33]), c_0_33]), c_0_15]), c_0_33])])).
cnf(c_0_100, negated_conjecture, (r2_lattice3(esk2_0,X1,esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)))|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk1_0,X1,esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk8_3(esk1_0,X1,esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X1))),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_96]), c_0_39])]), c_0_87]), c_0_97])).
cnf(c_0_101, negated_conjecture, (r2_lattice3(esk1_0,X1,esk7_3(esk2_0,X1,X2))|r2_lattice3(esk2_0,X1,esk6_3(esk2_0,X1,X2))|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,X2)|~m1_subset_1(esk7_3(esk2_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_98]), c_0_39])])).
cnf(c_0_102, negated_conjecture, (r2_lattice3(esk1_0,X1,esk7_3(esk2_0,X1,X2))|r1_yellow_0(esk2_0,X1)|m1_subset_1(esk6_3(esk2_0,X1,X2),u1_struct_0(esk1_0))|~r2_lattice3(esk2_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_99]), c_0_39])]), c_0_68])).
cnf(c_0_103, negated_conjecture, (r2_lattice3(esk2_0,X1,esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)))|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk1_0,X1,esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)))|~r1_yellow_0(esk1_0,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_41]), c_0_39])]), c_0_87]), c_0_97])).
cnf(c_0_104, negated_conjecture, (r2_lattice3(esk1_0,X1,esk7_3(esk2_0,X1,X2))|r2_lattice3(X3,X1,esk6_3(esk2_0,X1,X2))|r1_yellow_0(esk2_0,X1)|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))|~r2_lattice3(esk2_0,X1,X2)|~l1_orders_2(X3)|~m1_subset_1(esk6_3(esk2_0,X1,X2),u1_struct_0(X3))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_101]), c_0_33]), c_0_18]), c_0_15]), c_0_33])]), c_0_69]), c_0_102])).
cnf(c_0_105, negated_conjecture, (r2_lattice3(X1,X2,esk6_3(esk2_0,X2,k1_yellow_0(esk1_0,X2)))|r1_yellow_0(esk2_0,X2)|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(esk1_0,X2,esk7_3(esk2_0,X2,k1_yellow_0(esk1_0,X2)))|~r1_yellow_0(esk1_0,X2)|~l1_orders_2(X1)|~m1_subset_1(esk6_3(esk2_0,X2,k1_yellow_0(esk1_0,X2)),u1_struct_0(esk1_0))|~m1_subset_1(esk6_3(esk2_0,X2,k1_yellow_0(esk1_0,X2)),u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_103]), c_0_33]), c_0_18]), c_0_15]), c_0_33])])).
cnf(c_0_106, negated_conjecture, (r2_lattice3(esk1_0,X1,esk6_3(esk2_0,X1,X2))|r2_lattice3(esk1_0,X1,esk7_3(esk2_0,X1,X2))|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_104]), c_0_39])]), c_0_102])).
cnf(c_0_107, plain, (r1_orders_2(X2,esk7_3(X2,X3,X4),X1)|r1_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r1_orders_2(X2,X4,esk6_3(X2,X3,X4))|~r2_lattice3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_108, negated_conjecture, (r2_lattice3(esk1_0,X1,esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)))|r2_lattice3(X2,X1,esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)))|r1_yellow_0(esk2_0,X1)|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))|~r2_lattice3(esk2_0,X1,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~l1_orders_2(X2)|~m1_subset_1(esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)),u1_struct_0(X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_82])).
cnf(c_0_109, negated_conjecture, (r1_orders_2(esk2_0,esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X2)),X3)|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk1_0,X2,esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X2)))|~r2_lattice3(esk2_0,X1,k1_yellow_0(esk1_0,X2))|~r2_lattice3(esk2_0,X1,X3)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X2)),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_46]), c_0_15]), c_0_33]), c_0_33])]), c_0_82])).
cnf(c_0_110, negated_conjecture, (r2_lattice3(esk1_0,X1,esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)))|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_108]), c_0_39])])).
cnf(c_0_111, plain, (r2_lattice3(X1,X2,esk7_3(X1,X2,X3))|r1_yellow_0(X1,X2)|~r1_orders_2(X1,X3,esk6_3(X1,X2,X3))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_112, plain, (r1_orders_2(X2,esk7_3(X2,X3,X4),X1)|m1_subset_1(esk6_3(X2,X3,X4),u1_struct_0(X2))|r1_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r2_lattice3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_113, negated_conjecture, (r1_orders_2(esk2_0,esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)),X2)|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,k1_yellow_0(esk1_0,X1))|~r2_lattice3(esk2_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_109, c_0_110])).
cnf(c_0_114, negated_conjecture, (r2_lattice3(esk2_0,X1,esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X2)))|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk1_0,X2,esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X2)))|~r2_lattice3(esk2_0,X1,k1_yellow_0(esk1_0,X2))|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X2)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_46]), c_0_15]), c_0_33])]), c_0_82])).
cnf(c_0_115, plain, (r1_orders_2(X1,esk7_3(X2,X3,X4),X5)|r1_yellow_0(X2,X3)|m1_subset_1(esk6_3(X2,X3,X4),u1_struct_0(X2))|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(X2,X3,X4)|~r2_lattice3(X2,X3,X5)|~l1_orders_2(X1)|~l1_orders_2(X2)|~m1_subset_1(esk7_3(X2,X3,X4),u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_112]), c_0_57])).
cnf(c_0_116, negated_conjecture, (r1_orders_2(X1,esk4_2(esk2_0,X2),X3)|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(esk2_0,X2,X3)|~r1_yellow_0(esk2_0,X2)|~l1_orders_2(X1)|~m1_subset_1(esk4_2(esk2_0,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_18]), c_0_33]), c_0_15]), c_0_33])])).
cnf(c_0_117, negated_conjecture, (m1_subset_1(esk4_2(esk2_0,X1),u1_struct_0(esk1_0))|~r1_yellow_0(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_33]), c_0_15])])).
cnf(c_0_118, negated_conjecture, (r1_orders_2(X1,esk7_3(esk2_0,X2,k1_yellow_0(esk1_0,X2)),X3)|r1_yellow_0(esk2_0,X2)|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(esk2_0,X2,k1_yellow_0(esk1_0,X2))|~r2_lattice3(esk2_0,X2,X3)|~r1_yellow_0(esk1_0,X2)|~l1_orders_2(X1)|~m1_subset_1(esk6_3(esk2_0,X2,k1_yellow_0(esk1_0,X2)),u1_struct_0(esk1_0))|~m1_subset_1(esk7_3(esk2_0,X2,k1_yellow_0(esk1_0,X2)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_113]), c_0_33]), c_0_18]), c_0_15]), c_0_33]), c_0_33])]), c_0_87])).
cnf(c_0_119, negated_conjecture, (r2_lattice3(esk2_0,X1,esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)))|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_114, c_0_110])).
cnf(c_0_120, negated_conjecture, (r1_orders_2(X1,esk7_3(esk2_0,X2,X3),X4)|r1_yellow_0(esk2_0,X2)|m1_subset_1(esk6_3(esk2_0,X2,X3),u1_struct_0(esk1_0))|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(esk2_0,X2,X3)|~r2_lattice3(esk2_0,X2,X4)|~l1_orders_2(X1)|~m1_subset_1(esk7_3(esk2_0,X2,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_18]), c_0_33]), c_0_33]), c_0_15]), c_0_33]), c_0_33])])).
cnf(c_0_121, plain, (X1=esk4_2(X2,X3)|r2_lattice3(X4,X3,esk5_3(X2,X3,X1))|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))|~r2_lattice3(X2,X3,X1)|~r1_yellow_0(X2,X3)|~l1_orders_2(X4)|~l1_orders_2(X2)|~m1_subset_1(esk5_3(X2,X3,X1),u1_struct_0(X4))|~m1_subset_1(X1,u1_struct_0(X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_59]), c_0_64])).
cnf(c_0_122, negated_conjecture, (r1_orders_2(esk1_0,esk4_2(esk2_0,X1),X2)|~r2_lattice3(esk2_0,X1,X2)|~r1_yellow_0(esk2_0,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_116]), c_0_39])]), c_0_117])).
cnf(c_0_123, negated_conjecture, (r2_lattice3(X1,X2,esk4_2(esk2_0,X2))|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_yellow_0(esk2_0,X2)|~l1_orders_2(X1)|~m1_subset_1(esk4_2(esk2_0,X2),u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_18]), c_0_33]), c_0_15])])).
cnf(c_0_124, negated_conjecture, (r1_orders_2(esk1_0,esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)),X2)|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,k1_yellow_0(esk1_0,X1))|~r2_lattice3(esk2_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_118]), c_0_39])]), c_0_87])).
cnf(c_0_125, negated_conjecture, (r2_lattice3(X1,X2,esk7_3(esk2_0,X2,k1_yellow_0(esk1_0,X2)))|r1_yellow_0(esk2_0,X2)|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(esk2_0,X2,k1_yellow_0(esk1_0,X2))|~r1_yellow_0(esk1_0,X2)|~l1_orders_2(X1)|~m1_subset_1(esk6_3(esk2_0,X2,k1_yellow_0(esk1_0,X2)),u1_struct_0(esk1_0))|~m1_subset_1(esk7_3(esk2_0,X2,k1_yellow_0(esk1_0,X2)),u1_struct_0(X1))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_119]), c_0_33]), c_0_18]), c_0_15]), c_0_33])]), c_0_87])).
cnf(c_0_126, negated_conjecture, (r1_orders_2(esk1_0,esk7_3(esk2_0,X1,X2),X3)|r1_yellow_0(esk2_0,X1)|m1_subset_1(esk6_3(esk2_0,X1,X2),u1_struct_0(esk1_0))|~r2_lattice3(esk2_0,X1,X2)|~r2_lattice3(esk2_0,X1,X3)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_120]), c_0_39])]), c_0_68])).
cnf(c_0_127, negated_conjecture, (X1=esk4_2(X2,X3)|r2_lattice3(esk2_0,X3,esk5_3(X2,X3,X1))|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~r2_lattice3(X2,X3,X1)|~r1_yellow_0(X2,X3)|~l1_orders_2(X2)|~m1_subset_1(esk5_3(X2,X3,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_18]), c_0_33]), c_0_15]), c_0_33])])).
cnf(c_0_128, negated_conjecture, (r1_orders_2(esk1_0,k1_yellow_0(esk2_0,X1),X2)|~r2_lattice3(esk2_0,X1,X2)|~r1_yellow_0(esk2_0,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_44]), c_0_15])])).
cnf(c_0_129, negated_conjecture, (m1_subset_1(k1_yellow_0(esk2_0,X1),u1_struct_0(esk1_0))|~r1_yellow_0(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_44]), c_0_15])])).
cnf(c_0_130, negated_conjecture, (r2_lattice3(esk1_0,X1,esk4_2(esk2_0,X1))|~r1_yellow_0(esk2_0,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_123]), c_0_39])]), c_0_117])).
cnf(c_0_131, negated_conjecture, (esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X1))=k1_yellow_0(esk1_0,X2)|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,esk8_3(esk1_0,X2,esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X1))))|~r2_lattice3(esk1_0,X2,esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)))|~r2_lattice3(esk2_0,X1,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk8_3(esk1_0,X2,esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X1))),u1_struct_0(esk1_0))|~m1_subset_1(esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_124]), c_0_39])]), c_0_87])).
cnf(c_0_132, negated_conjecture, (r2_lattice3(esk1_0,X1,esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)))|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_125]), c_0_39])]), c_0_87])).
cnf(c_0_133, negated_conjecture, (esk7_3(esk2_0,X1,X2)=esk4_2(esk1_0,X3)|r1_yellow_0(esk2_0,X1)|m1_subset_1(esk6_3(esk2_0,X1,X2),u1_struct_0(esk1_0))|~r2_lattice3(esk2_0,X1,esk5_3(esk1_0,X3,esk7_3(esk2_0,X1,X2)))|~r2_lattice3(esk1_0,X3,esk7_3(esk2_0,X1,X2))|~r2_lattice3(esk2_0,X1,X2)|~r1_yellow_0(esk1_0,X3)|~m1_subset_1(esk5_3(esk1_0,X3,esk7_3(esk2_0,X1,X2)),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_126]), c_0_39])]), c_0_68])).
cnf(c_0_134, negated_conjecture, (X1=esk4_2(esk1_0,X2)|r2_lattice3(esk2_0,X2,esk5_3(esk1_0,X2,X1))|~r2_lattice3(esk1_0,X2,X1)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(esk5_3(esk1_0,X2,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_127]), c_0_39])])).
cnf(c_0_135, negated_conjecture, (k1_yellow_0(esk2_0,X1)=esk4_2(esk1_0,X2)|~r2_lattice3(esk2_0,X1,esk5_3(esk1_0,X2,k1_yellow_0(esk2_0,X1)))|~r2_lattice3(esk1_0,X2,k1_yellow_0(esk2_0,X1))|~r1_yellow_0(esk1_0,X2)|~r1_yellow_0(esk2_0,X1)|~m1_subset_1(esk5_3(esk1_0,X2,k1_yellow_0(esk2_0,X1)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_128]), c_0_39])]), c_0_129])).
cnf(c_0_136, negated_conjecture, (r2_lattice3(esk1_0,X1,k1_yellow_0(esk2_0,X1))|~r1_yellow_0(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_44]), c_0_15])])).
cnf(c_0_137, negated_conjecture, (r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk8_3(esk1_0,X1,esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X1))),u1_struct_0(esk1_0))|~m1_subset_1(esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_80]), c_0_87]), c_0_132]), c_0_97])).
cnf(c_0_138, negated_conjecture, (esk7_3(esk2_0,X1,X2)=esk4_2(esk1_0,X1)|r1_yellow_0(esk2_0,X1)|m1_subset_1(esk6_3(esk2_0,X1,X2),u1_struct_0(esk1_0))|~r2_lattice3(esk2_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk5_3(esk1_0,X1,esk7_3(esk2_0,X1,X2)),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134]), c_0_68]), c_0_102])).
cnf(c_0_139, negated_conjecture, (esk4_2(esk1_0,X1)=k1_yellow_0(esk2_0,X1)|~r1_yellow_0(esk1_0,X1)|~r1_yellow_0(esk2_0,X1)|~m1_subset_1(esk5_3(esk1_0,X1,k1_yellow_0(esk2_0,X1)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_134]), c_0_129]), c_0_136])).
cnf(c_0_140, negated_conjecture, (r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk6_3(esk2_0,X1,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_41]), c_0_39])]), c_0_87]), c_0_132]), c_0_97])).
cnf(c_0_141, negated_conjecture, (esk7_3(esk2_0,X1,X2)=esk4_2(esk1_0,X1)|r1_yellow_0(esk2_0,X1)|m1_subset_1(esk6_3(esk2_0,X1,X2),u1_struct_0(esk1_0))|~r2_lattice3(esk2_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_64]), c_0_39])]), c_0_68]), c_0_102])).
cnf(c_0_142, negated_conjecture, (esk4_2(esk1_0,X1)=k1_yellow_0(esk2_0,X1)|~r1_yellow_0(esk1_0,X1)|~r1_yellow_0(esk2_0,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_64]), c_0_39])]), c_0_129]), c_0_136])).
cnf(c_0_143, negated_conjecture, (esk7_3(esk2_0,X1,k1_yellow_0(esk1_0,X1))=esk4_2(esk1_0,X1)|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_141]), c_0_82])).
cnf(c_0_144, negated_conjecture, (k1_yellow_0(esk1_0,esk3_0)!=k1_yellow_0(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_145, negated_conjecture, (k1_yellow_0(esk2_0,X1)=k1_yellow_0(esk1_0,X1)|~r1_yellow_0(esk1_0,X1)|~r1_yellow_0(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_142]), c_0_39])])).
cnf(c_0_146, negated_conjecture, (r1_yellow_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_147, negated_conjecture, (r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_143]), c_0_77])).
cnf(c_0_148, negated_conjecture, (~r1_yellow_0(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_145]), c_0_146])])).
cnf(c_0_149, negated_conjecture, (r1_yellow_0(esk2_0,X1)|~r1_yellow_0(esk1_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_49]), c_0_82])).
cnf(c_0_150, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_149]), c_0_146])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 151
# Proof object clause steps            : 136
# Proof object formula steps           : 15
# Proof object conjectures             : 97
# Proof object clause conjectures      : 94
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 102
# Proof object simplifying inferences  : 269
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 35
# Processed clauses                    : 2232
# ...of these trivial                  : 71
# ...subsumed                          : 1171
# ...remaining for further processing  : 990
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 467
# Backward-rewritten                   : 5
# Generated clauses                    : 4183
# ...of the previous two non-trivial   : 3301
# Contextual simplify-reflections      : 349
# Paramodulations                      : 4094
# Factorizations                       : 0
# Equation resolutions                 : 91
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 477
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 469
# Current number of unprocessed clauses: 954
# ...number of literals in the above   : 10731
# Current number of archived formulas  : 0
# Current number of archived clauses   : 507
# Clause-clause subsumption calls (NU) : 469457
# Rec. Clause-clause subsumption calls : 7234
# Non-unit clause-clause subsumptions  : 1987
# Unit Clause-clause subsumption calls : 150
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 227643

# -------------------------------------------------
# User time                : 0.303 s
# System time              : 0.003 s
# Total time               : 0.306 s
# Maximum resident set size: 28788 pages
