# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t46_yellow_6, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~((~(r2_hidden(X3,k10_yellow_6(X1,X2)))&![X4]:(m2_yellow_6(X4,X1,X2)=>?[X5]:(m2_yellow_6(X5,X1,X4)&r2_hidden(X3,k10_yellow_6(X1,X5))))))))), file('data/MPTP/proofs/4babc7a1774942fd89daa2d636a45ee2.E_input', t46_yellow_6)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/4babc7a1774942fd89daa2d636a45ee2.E_input', dt_l1_pre_topc)).
fof(t27_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m2_yellow_6(X3,X1,X2)=>![X4]:(r2_waybel_0(X1,X3,X4)=>r2_waybel_0(X1,X2,X4))))), file('data/MPTP/proofs/4babc7a1774942fd89daa2d636a45ee2.E_input', t27_yellow_6)).
fof(dt_m2_yellow_6, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&(~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1)))))=>![X3]:(m2_yellow_6(X3,X1,X2)=>(~(v2_struct_0(X3))&(v4_orders_2(X3)&(v7_waybel_0(X3)&l1_waybel_0(X3,X1)))))), file('data/MPTP/proofs/4babc7a1774942fd89daa2d636a45ee2.E_input', dt_m2_yellow_6)).
fof(t9_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r1_waybel_0(X1,X2,X3)<=>~(r2_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3)))))), file('data/MPTP/proofs/4babc7a1774942fd89daa2d636a45ee2.E_input', t9_waybel_0)).
fof(t31_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(r2_waybel_0(X1,X2,X3)=>m2_yellow_6(k5_yellow_6(X1,X2,X3),X1,X2)))), file('data/MPTP/proofs/4babc7a1774942fd89daa2d636a45ee2.E_input', t31_yellow_6)).
fof(t30_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(r2_waybel_0(X1,X2,X3)=>(~(v2_struct_0(k5_yellow_6(X1,X2,X3)))&v7_waybel_0(k5_yellow_6(X1,X2,X3)))))), file('data/MPTP/proofs/4babc7a1774942fd89daa2d636a45ee2.E_input', t30_yellow_6)).
fof(d18_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k10_yellow_6(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>![X5]:(m1_connsp_2(X5,X1,X4)=>r1_waybel_0(X1,X2,X5)))))))), file('data/MPTP/proofs/4babc7a1774942fd89daa2d636a45ee2.E_input', d18_yellow_6)).
fof(dt_k10_yellow_6, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&(~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1)))))=>m1_subset_1(k10_yellow_6(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/4babc7a1774942fd89daa2d636a45ee2.E_input', dt_k10_yellow_6)).
fof(fc17_yellow_6, axiom, ![X1, X2, X3]:((l1_struct_0(X1)&(v4_orders_2(X2)&l1_waybel_0(X2,X1)))=>(v4_orders_2(k5_yellow_6(X1,X2,X3))&(v6_waybel_0(k5_yellow_6(X1,X2,X3),X1)&v2_yellow_6(k5_yellow_6(X1,X2,X3),X1,X2)))), file('data/MPTP/proofs/4babc7a1774942fd89daa2d636a45ee2.E_input', fc17_yellow_6)).
fof(t28_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(r1_waybel_0(X1,X2,X3)=>r2_waybel_0(X1,X2,X3)))), file('data/MPTP/proofs/4babc7a1774942fd89daa2d636a45ee2.E_input', t28_yellow_6)).
fof(t32_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3, X4]:(m2_yellow_6(X4,X1,X2)=>(X4=k5_yellow_6(X1,X2,X3)=>r1_waybel_0(X1,X4,X3))))), file('data/MPTP/proofs/4babc7a1774942fd89daa2d636a45ee2.E_input', t32_yellow_6)).
fof(c_0_12, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~((~(r2_hidden(X3,k10_yellow_6(X1,X2)))&![X4]:(m2_yellow_6(X4,X1,X2)=>?[X5]:(m2_yellow_6(X5,X1,X4)&r2_hidden(X3,k10_yellow_6(X1,X5)))))))))), inference(assume_negation,[status(cth)],[t46_yellow_6])).
fof(c_0_13, plain, ![X42]:(~l1_pre_topc(X42)|l1_struct_0(X42)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_14, negated_conjecture, ![X9]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&(v4_orders_2(esk2_0)&(v7_waybel_0(esk2_0)&l1_waybel_0(esk2_0,esk1_0))))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(~r2_hidden(esk3_0,k10_yellow_6(esk1_0,esk2_0))&((m2_yellow_6(esk4_1(X9),esk1_0,X9)|~m2_yellow_6(X9,esk1_0,esk2_0))&(r2_hidden(esk3_0,k10_yellow_6(esk1_0,esk4_1(X9)))|~m2_yellow_6(X9,esk1_0,esk2_0))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])])])).
fof(c_0_15, plain, ![X31, X32, X33, X34]:(v2_struct_0(X31)|~l1_struct_0(X31)|(v2_struct_0(X32)|(~v4_orders_2(X32)|(~v7_waybel_0(X32)|~l1_waybel_0(X32,X31)))|(~m2_yellow_6(X33,X31,X32)|(~r2_waybel_0(X31,X33,X34)|r2_waybel_0(X31,X32,X34))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t27_yellow_6])])])])).
cnf(c_0_16, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, plain, ![X19, X20, X21]:((~v2_struct_0(X21)|~m2_yellow_6(X21,X19,X20)|(v2_struct_0(X19)|~l1_struct_0(X19)|(v2_struct_0(X20)|(~v4_orders_2(X20)|(~v7_waybel_0(X20)|~l1_waybel_0(X20,X19))))))&((v4_orders_2(X21)|~m2_yellow_6(X21,X19,X20)|(v2_struct_0(X19)|~l1_struct_0(X19)|(v2_struct_0(X20)|(~v4_orders_2(X20)|(~v7_waybel_0(X20)|~l1_waybel_0(X20,X19))))))&((v7_waybel_0(X21)|~m2_yellow_6(X21,X19,X20)|(v2_struct_0(X19)|~l1_struct_0(X19)|(v2_struct_0(X20)|(~v4_orders_2(X20)|(~v7_waybel_0(X20)|~l1_waybel_0(X20,X19))))))&(l1_waybel_0(X21,X19)|~m2_yellow_6(X21,X19,X20)|(v2_struct_0(X19)|~l1_struct_0(X19)|(v2_struct_0(X20)|(~v4_orders_2(X20)|(~v7_waybel_0(X20)|~l1_waybel_0(X20,X19))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_yellow_6])])])])])).
cnf(c_0_19, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r2_waybel_0(X1,X2,X4)|~l1_struct_0(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~m2_yellow_6(X3,X1,X2)|~r2_waybel_0(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, negated_conjecture, (m2_yellow_6(esk4_1(X1),esk1_0,X1)|~m2_yellow_6(X1,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, negated_conjecture, (l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_22, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_23, plain, ![X39, X40, X41]:((~r1_waybel_0(X39,X40,X41)|~r2_waybel_0(X39,X40,k6_subset_1(u1_struct_0(X39),X41))|(v2_struct_0(X40)|~l1_waybel_0(X40,X39))|(v2_struct_0(X39)|~l1_struct_0(X39)))&(r2_waybel_0(X39,X40,k6_subset_1(u1_struct_0(X39),X41))|r1_waybel_0(X39,X40,X41)|(v2_struct_0(X40)|~l1_waybel_0(X40,X39))|(v2_struct_0(X39)|~l1_struct_0(X39)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_waybel_0])])])])])).
cnf(c_0_24, plain, (l1_waybel_0(X1,X2)|v2_struct_0(X2)|v2_struct_0(X3)|~m2_yellow_6(X1,X2,X3)|~l1_struct_0(X2)|~v4_orders_2(X3)|~v7_waybel_0(X3)|~l1_waybel_0(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (v2_struct_0(X2)|v2_struct_0(X3)|~v2_struct_0(X1)|~m2_yellow_6(X1,X2,X3)|~l1_struct_0(X2)|~v4_orders_2(X3)|~v7_waybel_0(X3)|~l1_waybel_0(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, negated_conjecture, (r2_waybel_0(esk1_0,X1,X2)|v2_struct_0(X1)|~r2_waybel_0(esk1_0,esk4_1(X1),X2)|~m2_yellow_6(X1,esk1_0,esk2_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_27, plain, (r2_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3))|r1_waybel_0(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28, negated_conjecture, (l1_waybel_0(esk4_1(X1),esk1_0)|v2_struct_0(X1)|~m2_yellow_6(X1,esk1_0,esk2_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_29, negated_conjecture, (v2_struct_0(X1)|~m2_yellow_6(X1,esk1_0,esk2_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~v2_struct_0(esk4_1(X1))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_20]), c_0_21])]), c_0_22])).
fof(c_0_30, plain, ![X16, X17, X18]:(v2_struct_0(X16)|~l1_struct_0(X16)|(v2_struct_0(X17)|(~v4_orders_2(X17)|(~v7_waybel_0(X17)|~l1_waybel_0(X17,X16)))|(~r2_waybel_0(X16,X17,X18)|m2_yellow_6(k5_yellow_6(X16,X17,X18),X16,X17)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t31_yellow_6])])])])).
cnf(c_0_31, negated_conjecture, (r2_waybel_0(esk1_0,X1,k6_subset_1(u1_struct_0(esk1_0),X2))|r1_waybel_0(esk1_0,esk4_1(X1),X2)|v2_struct_0(X1)|~m2_yellow_6(X1,esk1_0,esk2_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_21])]), c_0_22]), c_0_28]), c_0_29])).
cnf(c_0_32, plain, (v2_struct_0(X1)|v2_struct_0(X2)|m2_yellow_6(k5_yellow_6(X1,X2,X3),X1,X2)|~l1_struct_0(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~r2_waybel_0(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_33, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_34, negated_conjecture, (v7_waybel_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_35, negated_conjecture, (v4_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_36, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_37, negated_conjecture, (r2_waybel_0(esk1_0,k5_yellow_6(esk1_0,esk2_0,X1),k6_subset_1(u1_struct_0(esk1_0),X2))|r1_waybel_0(esk1_0,esk4_1(k5_yellow_6(esk1_0,esk2_0,X1)),X2)|v2_struct_0(k5_yellow_6(esk1_0,esk2_0,X1))|~r2_waybel_0(esk1_0,esk2_0,X1)|~l1_waybel_0(k5_yellow_6(esk1_0,esk2_0,X1),esk1_0)|~v7_waybel_0(k5_yellow_6(esk1_0,esk2_0,X1))|~v4_orders_2(k5_yellow_6(esk1_0,esk2_0,X1))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_21]), c_0_33]), c_0_34]), c_0_35])]), c_0_36]), c_0_22])).
cnf(c_0_38, plain, (l1_waybel_0(k5_yellow_6(X1,X2,X3),X1)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X1,X2,X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)), inference(spm,[status(thm)],[c_0_24, c_0_32])).
fof(c_0_39, plain, ![X46, X47, X48]:((~v2_struct_0(k5_yellow_6(X46,X47,X48))|~r2_waybel_0(X46,X47,X48)|(v2_struct_0(X47)|(~v4_orders_2(X47)|(~v7_waybel_0(X47)|~l1_waybel_0(X47,X46))))|(v2_struct_0(X46)|~l1_struct_0(X46)))&(v7_waybel_0(k5_yellow_6(X46,X47,X48))|~r2_waybel_0(X46,X47,X48)|(v2_struct_0(X47)|(~v4_orders_2(X47)|(~v7_waybel_0(X47)|~l1_waybel_0(X47,X46))))|(v2_struct_0(X46)|~l1_struct_0(X46)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t30_yellow_6])])])])])).
fof(c_0_40, plain, ![X22, X23, X24, X25, X26, X30]:(((~r2_hidden(X25,X24)|(~m1_connsp_2(X26,X22,X25)|r1_waybel_0(X22,X23,X26))|~m1_subset_1(X25,u1_struct_0(X22))|X24!=k10_yellow_6(X22,X23)|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X22)))|(v2_struct_0(X23)|(~v4_orders_2(X23)|(~v7_waybel_0(X23)|~l1_waybel_0(X23,X22))))|(v2_struct_0(X22)|(~v2_pre_topc(X22)|~l1_pre_topc(X22))))&((m1_connsp_2(esk5_4(X22,X23,X24,X25),X22,X25)|r2_hidden(X25,X24)|~m1_subset_1(X25,u1_struct_0(X22))|X24!=k10_yellow_6(X22,X23)|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X22)))|(v2_struct_0(X23)|(~v4_orders_2(X23)|(~v7_waybel_0(X23)|~l1_waybel_0(X23,X22))))|(v2_struct_0(X22)|(~v2_pre_topc(X22)|~l1_pre_topc(X22))))&(~r1_waybel_0(X22,X23,esk5_4(X22,X23,X24,X25))|r2_hidden(X25,X24)|~m1_subset_1(X25,u1_struct_0(X22))|X24!=k10_yellow_6(X22,X23)|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X22)))|(v2_struct_0(X23)|(~v4_orders_2(X23)|(~v7_waybel_0(X23)|~l1_waybel_0(X23,X22))))|(v2_struct_0(X22)|(~v2_pre_topc(X22)|~l1_pre_topc(X22))))))&((m1_subset_1(esk6_3(X22,X23,X24),u1_struct_0(X22))|X24=k10_yellow_6(X22,X23)|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X22)))|(v2_struct_0(X23)|(~v4_orders_2(X23)|(~v7_waybel_0(X23)|~l1_waybel_0(X23,X22))))|(v2_struct_0(X22)|(~v2_pre_topc(X22)|~l1_pre_topc(X22))))&(((m1_connsp_2(esk7_3(X22,X23,X24),X22,esk6_3(X22,X23,X24))|~r2_hidden(esk6_3(X22,X23,X24),X24)|X24=k10_yellow_6(X22,X23)|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X22)))|(v2_struct_0(X23)|(~v4_orders_2(X23)|(~v7_waybel_0(X23)|~l1_waybel_0(X23,X22))))|(v2_struct_0(X22)|(~v2_pre_topc(X22)|~l1_pre_topc(X22))))&(~r1_waybel_0(X22,X23,esk7_3(X22,X23,X24))|~r2_hidden(esk6_3(X22,X23,X24),X24)|X24=k10_yellow_6(X22,X23)|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X22)))|(v2_struct_0(X23)|(~v4_orders_2(X23)|(~v7_waybel_0(X23)|~l1_waybel_0(X23,X22))))|(v2_struct_0(X22)|(~v2_pre_topc(X22)|~l1_pre_topc(X22)))))&(r2_hidden(esk6_3(X22,X23,X24),X24)|(~m1_connsp_2(X30,X22,esk6_3(X22,X23,X24))|r1_waybel_0(X22,X23,X30))|X24=k10_yellow_6(X22,X23)|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X22)))|(v2_struct_0(X23)|(~v4_orders_2(X23)|(~v7_waybel_0(X23)|~l1_waybel_0(X23,X22))))|(v2_struct_0(X22)|(~v2_pre_topc(X22)|~l1_pre_topc(X22))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d18_yellow_6])])])])])])).
fof(c_0_41, plain, ![X11, X12]:(v2_struct_0(X11)|(~v2_pre_topc(X11)|~l1_pre_topc(X11))|(v2_struct_0(X12)|(~v4_orders_2(X12)|(~v7_waybel_0(X12)|~l1_waybel_0(X12,X11))))|m1_subset_1(k10_yellow_6(X11,X12),k1_zfmisc_1(u1_struct_0(X11)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k10_yellow_6])])])).
cnf(c_0_42, negated_conjecture, (r2_waybel_0(esk1_0,k5_yellow_6(esk1_0,esk2_0,X1),k6_subset_1(u1_struct_0(esk1_0),X2))|r1_waybel_0(esk1_0,esk4_1(k5_yellow_6(esk1_0,esk2_0,X1)),X2)|v2_struct_0(k5_yellow_6(esk1_0,esk2_0,X1))|~r2_waybel_0(esk1_0,esk2_0,X1)|~v7_waybel_0(k5_yellow_6(esk1_0,esk2_0,X1))|~v4_orders_2(k5_yellow_6(esk1_0,esk2_0,X1))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_21]), c_0_33]), c_0_34]), c_0_35])]), c_0_22]), c_0_36])).
cnf(c_0_43, plain, (v7_waybel_0(k5_yellow_6(X1,X2,X3))|v2_struct_0(X2)|v2_struct_0(X1)|~r2_waybel_0(X1,X2,X3)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_44, plain, ![X43, X44, X45]:((v4_orders_2(k5_yellow_6(X43,X44,X45))|(~l1_struct_0(X43)|(~v4_orders_2(X44)|~l1_waybel_0(X44,X43))))&((v6_waybel_0(k5_yellow_6(X43,X44,X45),X43)|(~l1_struct_0(X43)|(~v4_orders_2(X44)|~l1_waybel_0(X44,X43))))&(v2_yellow_6(k5_yellow_6(X43,X44,X45),X43,X44)|(~l1_struct_0(X43)|(~v4_orders_2(X44)|~l1_waybel_0(X44,X43)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc17_yellow_6])])])).
cnf(c_0_45, plain, (r1_waybel_0(X4,X5,X3)|v2_struct_0(X5)|v2_struct_0(X4)|~r2_hidden(X1,X2)|~m1_connsp_2(X3,X4,X1)|~m1_subset_1(X1,u1_struct_0(X4))|X2!=k10_yellow_6(X4,X5)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X4)))|~v4_orders_2(X5)|~v7_waybel_0(X5)|~l1_waybel_0(X5,X4)|~v2_pre_topc(X4)|~l1_pre_topc(X4)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_46, plain, (v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(k10_yellow_6(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_47, plain, ![X13, X14, X15]:(v2_struct_0(X13)|~l1_struct_0(X13)|(v2_struct_0(X14)|(~v4_orders_2(X14)|(~v7_waybel_0(X14)|~l1_waybel_0(X14,X13)))|(~r1_waybel_0(X13,X14,X15)|r2_waybel_0(X13,X14,X15)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t28_yellow_6])])])])).
cnf(c_0_48, negated_conjecture, (r2_waybel_0(esk1_0,k5_yellow_6(esk1_0,esk2_0,X1),k6_subset_1(u1_struct_0(esk1_0),X2))|r1_waybel_0(esk1_0,esk4_1(k5_yellow_6(esk1_0,esk2_0,X1)),X2)|v2_struct_0(k5_yellow_6(esk1_0,esk2_0,X1))|~r2_waybel_0(esk1_0,esk2_0,X1)|~v4_orders_2(k5_yellow_6(esk1_0,esk2_0,X1))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_21]), c_0_33]), c_0_34]), c_0_35])]), c_0_36]), c_0_22])).
cnf(c_0_49, plain, (v4_orders_2(k5_yellow_6(X1,X2,X3))|~l1_struct_0(X1)|~v4_orders_2(X2)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_50, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X2)|k10_yellow_6(X1,X4)!=k10_yellow_6(X1,X2)|~m1_connsp_2(X3,X1,X5)|~r2_hidden(X5,k10_yellow_6(X1,X4))|~m1_subset_1(X5,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~l1_waybel_0(X4,X1)|~v7_waybel_0(X2)|~v7_waybel_0(X4)|~v4_orders_2(X2)|~v4_orders_2(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_51, negated_conjecture, (r2_hidden(esk3_0,k10_yellow_6(esk1_0,esk4_1(X1)))|~m2_yellow_6(X1,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_52, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_53, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_54, plain, (v2_struct_0(X2)|v2_struct_0(X1)|~r1_waybel_0(X1,X2,X3)|~r2_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_55, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r2_waybel_0(X1,X2,X3)|~l1_struct_0(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~r1_waybel_0(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_56, negated_conjecture, (r2_waybel_0(esk1_0,k5_yellow_6(esk1_0,esk2_0,X1),k6_subset_1(u1_struct_0(esk1_0),X2))|r1_waybel_0(esk1_0,esk4_1(k5_yellow_6(esk1_0,esk2_0,X1)),X2)|v2_struct_0(k5_yellow_6(esk1_0,esk2_0,X1))|~r2_waybel_0(esk1_0,esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_21]), c_0_33]), c_0_35])])).
cnf(c_0_57, negated_conjecture, (r1_waybel_0(esk1_0,X1,X2)|v2_struct_0(esk4_1(X3))|v2_struct_0(X1)|k10_yellow_6(esk1_0,esk4_1(X3))!=k10_yellow_6(esk1_0,X1)|~m1_connsp_2(X2,esk1_0,esk3_0)|~m2_yellow_6(X3,esk1_0,esk2_0)|~l1_waybel_0(esk4_1(X3),esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(esk4_1(X3))|~v7_waybel_0(X1)|~v4_orders_2(esk4_1(X3))|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52]), c_0_17]), c_0_53])]), c_0_22])).
cnf(c_0_58, plain, (m1_connsp_2(esk5_4(X1,X2,X3,X4),X1,X4)|r2_hidden(X4,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k10_yellow_6(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_59, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r1_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3))|~r1_waybel_0(X1,X2,X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_60, negated_conjecture, (r2_waybel_0(esk1_0,k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1)),k6_subset_1(u1_struct_0(esk1_0),X2))|r1_waybel_0(esk1_0,esk4_1(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1))),X2)|r1_waybel_0(esk1_0,esk2_0,X1)|v2_struct_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_27]), c_0_21]), c_0_33])]), c_0_36]), c_0_22])).
cnf(c_0_61, negated_conjecture, (r1_waybel_0(esk1_0,X1,esk5_4(esk1_0,X2,X3,esk3_0))|r2_hidden(esk3_0,X3)|v2_struct_0(esk4_1(X4))|v2_struct_0(X2)|v2_struct_0(X1)|k10_yellow_6(esk1_0,esk4_1(X4))!=k10_yellow_6(esk1_0,X1)|X3!=k10_yellow_6(esk1_0,X2)|~m2_yellow_6(X4,esk1_0,esk2_0)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(esk4_1(X4),esk1_0)|~l1_waybel_0(X1,esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(esk4_1(X4))|~v7_waybel_0(X1)|~v7_waybel_0(X2)|~v4_orders_2(esk4_1(X4))|~v4_orders_2(X1)|~v4_orders_2(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_52]), c_0_17]), c_0_53])]), c_0_22])).
cnf(c_0_62, negated_conjecture, (r2_waybel_0(esk1_0,k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1)),k6_subset_1(u1_struct_0(esk1_0),k6_subset_1(u1_struct_0(esk1_0),X2)))|r1_waybel_0(esk1_0,esk2_0,X1)|v2_struct_0(esk4_1(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1))))|v2_struct_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1)))|~r1_waybel_0(esk1_0,esk4_1(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1))),X2)|~l1_waybel_0(esk4_1(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1))),esk1_0)|~v7_waybel_0(esk4_1(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1))))|~v4_orders_2(esk4_1(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_21])]), c_0_22])).
cnf(c_0_63, negated_conjecture, (r1_waybel_0(esk1_0,esk4_1(X1),esk5_4(esk1_0,X2,X3,esk3_0))|r2_hidden(esk3_0,X3)|v2_struct_0(esk4_1(X1))|v2_struct_0(X2)|X3!=k10_yellow_6(esk1_0,X2)|~m2_yellow_6(X1,esk1_0,esk2_0)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(esk4_1(X1),esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(esk4_1(X1))|~v7_waybel_0(X2)|~v4_orders_2(esk4_1(X1))|~v4_orders_2(X2)), inference(er,[status(thm)],[c_0_61])).
cnf(c_0_64, plain, (v7_waybel_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m2_yellow_6(X1,X2,X3)|~l1_struct_0(X2)|~v4_orders_2(X3)|~v7_waybel_0(X3)|~l1_waybel_0(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_65, plain, (v4_orders_2(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m2_yellow_6(X1,X2,X3)|~l1_struct_0(X2)|~v4_orders_2(X3)|~v7_waybel_0(X3)|~l1_waybel_0(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_66, negated_conjecture, (r2_waybel_0(esk1_0,k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1)),k6_subset_1(u1_struct_0(esk1_0),k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X2,X3,esk3_0))))|r1_waybel_0(esk1_0,esk2_0,X1)|r2_hidden(esk3_0,X3)|v2_struct_0(esk4_1(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1))))|v2_struct_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1)))|v2_struct_0(X2)|X3!=k10_yellow_6(esk1_0,X2)|~m2_yellow_6(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1)),esk1_0,esk2_0)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(esk4_1(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1))),esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(esk4_1(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1))))|~v7_waybel_0(X2)|~v4_orders_2(esk4_1(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1))))|~v4_orders_2(X2)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_67, negated_conjecture, (v7_waybel_0(esk4_1(X1))|v2_struct_0(X1)|~m2_yellow_6(X1,esk1_0,esk2_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_68, negated_conjecture, (v4_orders_2(esk4_1(X1))|v2_struct_0(X1)|~m2_yellow_6(X1,esk1_0,esk2_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_20]), c_0_21])]), c_0_22])).
fof(c_0_69, plain, ![X35, X36, X37, X38]:(v2_struct_0(X35)|~l1_struct_0(X35)|(v2_struct_0(X36)|(~v4_orders_2(X36)|(~v7_waybel_0(X36)|~l1_waybel_0(X36,X35)))|(~m2_yellow_6(X38,X35,X36)|(X38!=k5_yellow_6(X35,X36,X37)|r1_waybel_0(X35,X38,X37))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t32_yellow_6])])])])).
cnf(c_0_70, negated_conjecture, (r2_waybel_0(esk1_0,k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1)),k6_subset_1(u1_struct_0(esk1_0),k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X2,X3,esk3_0))))|r1_waybel_0(esk1_0,esk2_0,X1)|r2_hidden(esk3_0,X3)|v2_struct_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1)))|v2_struct_0(X2)|X3!=k10_yellow_6(esk1_0,X2)|~m2_yellow_6(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1)),esk1_0,esk2_0)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1)),esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1)))|~v7_waybel_0(X2)|~v4_orders_2(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1)))|~v4_orders_2(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68]), c_0_28]), c_0_29])).
cnf(c_0_71, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r1_waybel_0(X1,X3,X4)|~l1_struct_0(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~m2_yellow_6(X3,X1,X2)|X3!=k5_yellow_6(X1,X2,X4)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_72, negated_conjecture, (r1_waybel_0(esk1_0,esk2_0,X1)|r2_hidden(esk3_0,X2)|v2_struct_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1)))|v2_struct_0(X3)|X2!=k10_yellow_6(esk1_0,X3)|~r1_waybel_0(esk1_0,k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1)),k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X3,X2,esk3_0)))|~m2_yellow_6(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1)),esk1_0,esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1)),esk1_0)|~l1_waybel_0(X3,esk1_0)|~v7_waybel_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1)))|~v7_waybel_0(X3)|~v4_orders_2(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),X1)))|~v4_orders_2(X3)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_70]), c_0_21])]), c_0_22])).
cnf(c_0_73, plain, (r1_waybel_0(X1,k5_yellow_6(X1,X2,X3),X3)|v2_struct_0(X1)|v2_struct_0(X2)|~l1_struct_0(X1)|~m2_yellow_6(k5_yellow_6(X1,X2,X3),X1,X2)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)), inference(er,[status(thm)],[c_0_71])).
cnf(c_0_74, negated_conjecture, (r1_waybel_0(esk1_0,esk2_0,esk5_4(esk1_0,X1,X2,esk3_0))|r2_hidden(esk3_0,X2)|v2_struct_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0))))|v2_struct_0(X1)|X2!=k10_yellow_6(esk1_0,X1)|~m2_yellow_6(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0))),esk1_0,esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0))),esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0))))|~v7_waybel_0(X1)|~v4_orders_2(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0))))|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_21]), c_0_33]), c_0_34]), c_0_35])]), c_0_22]), c_0_36])).
cnf(c_0_75, negated_conjecture, (r1_waybel_0(esk1_0,esk2_0,esk5_4(esk1_0,X1,X2,esk3_0))|r2_hidden(esk3_0,X2)|v2_struct_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0))))|v2_struct_0(X1)|X2!=k10_yellow_6(esk1_0,X1)|~r2_waybel_0(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0))),esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0))))|~v7_waybel_0(X1)|~v4_orders_2(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0))))|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_32]), c_0_21]), c_0_33]), c_0_34]), c_0_35])]), c_0_36]), c_0_22])).
cnf(c_0_76, negated_conjecture, (r1_waybel_0(esk1_0,esk2_0,esk5_4(esk1_0,X1,X2,esk3_0))|r2_hidden(esk3_0,X2)|v2_struct_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0))))|v2_struct_0(X1)|X2!=k10_yellow_6(esk1_0,X1)|~r2_waybel_0(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0))),esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0))))|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_43]), c_0_21]), c_0_33]), c_0_34]), c_0_35])]), c_0_36]), c_0_22])).
cnf(c_0_77, negated_conjecture, (r1_waybel_0(esk1_0,esk2_0,esk5_4(esk1_0,X1,X2,esk3_0))|r2_hidden(esk3_0,X2)|v2_struct_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0))))|v2_struct_0(X1)|X2!=k10_yellow_6(esk1_0,X1)|~r2_waybel_0(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0))),esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_49]), c_0_21]), c_0_33]), c_0_35])])).
cnf(c_0_78, plain, (v2_struct_0(X2)|v2_struct_0(X1)|~v2_struct_0(k5_yellow_6(X1,X2,X3))|~r2_waybel_0(X1,X2,X3)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_79, negated_conjecture, (r1_waybel_0(esk1_0,esk2_0,esk5_4(esk1_0,X1,X2,esk3_0))|r2_hidden(esk3_0,X2)|v2_struct_0(k5_yellow_6(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0))))|v2_struct_0(X1)|X2!=k10_yellow_6(esk1_0,X1)|~r2_waybel_0(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_38]), c_0_21]), c_0_33]), c_0_34]), c_0_35])]), c_0_22]), c_0_36])).
cnf(c_0_80, negated_conjecture, (r1_waybel_0(esk1_0,esk2_0,esk5_4(esk1_0,X1,X2,esk3_0))|r2_hidden(esk3_0,X2)|v2_struct_0(X1)|X2!=k10_yellow_6(esk1_0,X1)|~r2_waybel_0(esk1_0,esk2_0,k6_subset_1(u1_struct_0(esk1_0),esk5_4(esk1_0,X1,X2,esk3_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_21]), c_0_33]), c_0_34]), c_0_35])]), c_0_36]), c_0_22])).
cnf(c_0_81, plain, (r2_hidden(X4,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_waybel_0(X1,X2,esk5_4(X1,X2,X3,X4))|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k10_yellow_6(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_82, negated_conjecture, (r1_waybel_0(esk1_0,esk2_0,esk5_4(esk1_0,X1,X2,esk3_0))|r2_hidden(esk3_0,X2)|v2_struct_0(X1)|X2!=k10_yellow_6(esk1_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_27]), c_0_21]), c_0_33])]), c_0_36]), c_0_22])).
cnf(c_0_83, negated_conjecture, (~r2_hidden(esk3_0,k10_yellow_6(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_84, negated_conjecture, (r2_hidden(esk3_0,X1)|X1!=k10_yellow_6(esk1_0,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_52]), c_0_33]), c_0_34]), c_0_35]), c_0_17]), c_0_53])]), c_0_36]), c_0_22])).
cnf(c_0_85, negated_conjecture, (~m1_subset_1(k10_yellow_6(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_86, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_46]), c_0_33]), c_0_34]), c_0_35]), c_0_17]), c_0_53])]), c_0_36]), c_0_22]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 87
# Proof object clause steps            : 62
# Proof object formula steps           : 25
# Proof object conjectures             : 43
# Proof object clause conjectures      : 40
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 29
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 33
# Proof object simplifying inferences  : 130
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 35
# Processed clauses                    : 5735
# ...of these trivial                  : 0
# ...subsumed                          : 3817
# ...remaining for further processing  : 1918
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 616
# Backward-rewritten                   : 4
# Generated clauses                    : 50256
# ...of the previous two non-trivial   : 50005
# Contextual simplify-reflections      : 374
# Paramodulations                      : 50024
# Factorizations                       : 212
# Equation resolutions                 : 20
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1263
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 1249
# Current number of unprocessed clauses: 43542
# ...number of literals in the above   : 421659
# Current number of archived formulas  : 0
# Current number of archived clauses   : 655
# Clause-clause subsumption calls (NU) : 1668503
# Rec. Clause-clause subsumption calls : 59248
# Non-unit clause-clause subsumptions  : 4792
# Unit Clause-clause subsumption calls : 719
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 50
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3774977

# -------------------------------------------------
# User time                : 2.573 s
# System time              : 0.020 s
# Total time               : 2.593 s
# Maximum resident set size: 47704 pages
