# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', d1_xboole_0)).
fof(t103_zfmisc_1, axiom, ![X1, X2, X3, X4]:~((r1_tarski(X1,k2_zfmisc_1(X2,X3))&(r2_hidden(X4,X1)&![X5, X6]:~((r2_hidden(X5,X2)&(r2_hidden(X6,X3)&X4=k4_tarski(X5,X6))))))), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', t103_zfmisc_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', d3_tarski)).
fof(t71_mcart_1, conjecture, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X8))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k7_mcart_1(X1,X2,X3,X5)))))), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', t71_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', d3_zfmisc_1)).
fof(d4_zfmisc_1, axiom, ![X1, X2, X3, X4]:k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', d4_zfmisc_1)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', l13_xboole_0)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', d2_subset_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', t113_zfmisc_1)).
fof(t8_boole, axiom, ![X1, X2]:~((v1_xboole_0(X1)&(~(X1=X2)&v1_xboole_0(X2)))), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', t8_boole)).
fof(t55_mcart_1, axiom, ![X1, X2, X3, X4]:((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(X4=k1_xboole_0))))<=>~(k4_zfmisc_1(X1,X2,X3,X4)=k1_xboole_0)), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', t55_mcart_1)).
fof(rc1_xboole_0, axiom, ?[X1]:v1_xboole_0(X1), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', rc1_xboole_0)).
fof(t152_zfmisc_1, axiom, ![X1, X2]:~(r2_hidden(X1,k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', t152_zfmisc_1)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', fc1_subset_1)).
fof(t10_mcart_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(r2_hidden(k1_mcart_1(X1),X2)&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', t10_mcart_1)).
fof(t50_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))&(k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))&k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4))))))))), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', t50_mcart_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', d3_mcart_1)).
fof(t23_mcart_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r2_hidden(X1,X2)=>X1=k4_tarski(k1_mcart_1(X1),k2_mcart_1(X1)))), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', t23_mcart_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', fc6_relat_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/4be991e1e7c84bf38f4e3b950378509f.E_input', t1_subset)).
fof(c_0_20, plain, ![X9, X10, X11]:((~v1_xboole_0(X9)|~r2_hidden(X10,X9))&(r2_hidden(esk1_1(X11),X11)|v1_xboole_0(X11))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_21, plain, ![X13, X14, X15, X16]:((r2_hidden(esk2_4(X13,X14,X15,X16),X14)|~r2_hidden(X16,X13)|~r1_tarski(X13,k2_zfmisc_1(X14,X15)))&((r2_hidden(esk3_4(X13,X14,X15,X16),X15)|~r2_hidden(X16,X13)|~r1_tarski(X13,k2_zfmisc_1(X14,X15)))&(X16=k4_tarski(esk2_4(X13,X14,X15,X16),esk3_4(X13,X14,X15,X16))|~r2_hidden(X16,X13)|~r1_tarski(X13,k2_zfmisc_1(X14,X15))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t103_zfmisc_1])])])])])])).
fof(c_0_22, plain, ![X19, X20, X21, X22, X23]:((~r1_tarski(X19,X20)|(~r2_hidden(X21,X19)|r2_hidden(X21,X20)))&((r2_hidden(esk4_2(X22,X23),X22)|r1_tarski(X22,X23))&(~r2_hidden(esk4_2(X22,X23),X23)|r1_tarski(X22,X23)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_23, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24, plain, (r2_hidden(esk2_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X1)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_25, plain, (r1_tarski(X1,X2)|~r2_hidden(esk4_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_27, negated_conjecture, ~(![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X8))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k7_mcart_1(X1,X2,X3,X5))))))), inference(assume_negation,[status(cth)],[t71_mcart_1])).
cnf(c_0_28, plain, (~r1_tarski(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X4,X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_29, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_30, negated_conjecture, ![X30, X31, X32]:(m1_subset_1(esk9_0,k3_zfmisc_1(esk5_0,esk6_0,esk7_0))&((~m1_subset_1(X30,esk5_0)|(~m1_subset_1(X31,esk6_0)|(~m1_subset_1(X32,esk7_0)|(esk9_0!=k3_mcart_1(X30,X31,X32)|esk8_0=X32))))&(esk5_0!=k1_xboole_0&(esk6_0!=k1_xboole_0&(esk7_0!=k1_xboole_0&esk8_0!=k7_mcart_1(esk5_0,esk6_0,esk7_0,esk9_0)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])])])).
fof(c_0_31, plain, ![X33, X34, X35]:k3_zfmisc_1(X33,X34,X35)=k2_zfmisc_1(k2_zfmisc_1(X33,X34),X35), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_32, plain, ![X36, X37, X38, X39]:k4_zfmisc_1(X36,X37,X38,X39)=k2_zfmisc_1(k3_zfmisc_1(X36,X37,X38),X39), inference(variable_rename,[status(thm)],[d4_zfmisc_1])).
fof(c_0_33, plain, ![X40]:(~v1_xboole_0(X40)|X40=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_34, plain, (~r2_hidden(X1,k2_zfmisc_1(X2,X3))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35, plain, (r2_hidden(esk1_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_36, plain, ![X41, X42]:(((~m1_subset_1(X42,X41)|r2_hidden(X42,X41)|v1_xboole_0(X41))&(~r2_hidden(X42,X41)|m1_subset_1(X42,X41)|v1_xboole_0(X41)))&((~m1_subset_1(X42,X41)|v1_xboole_0(X42)|~v1_xboole_0(X41))&(~v1_xboole_0(X42)|m1_subset_1(X42,X41)|~v1_xboole_0(X41)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
cnf(c_0_37, negated_conjecture, (m1_subset_1(esk9_0,k3_zfmisc_1(esk5_0,esk6_0,esk7_0))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_39, plain, ![X59, X60]:((k2_zfmisc_1(X59,X60)!=k1_xboole_0|(X59=k1_xboole_0|X60=k1_xboole_0))&((X59!=k1_xboole_0|k2_zfmisc_1(X59,X60)=k1_xboole_0)&(X60!=k1_xboole_0|k2_zfmisc_1(X59,X60)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_40, plain, ![X61, X62]:(~v1_xboole_0(X61)|(X61=X62|~v1_xboole_0(X62))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_boole])])])).
fof(c_0_41, plain, ![X43, X44, X45, X46]:((X43=k1_xboole_0|(X44=k1_xboole_0|(X45=k1_xboole_0|X46=k1_xboole_0))|k4_zfmisc_1(X43,X44,X45,X46)!=k1_xboole_0)&((X43!=k1_xboole_0|k4_zfmisc_1(X43,X44,X45,X46)=k1_xboole_0)&((X44!=k1_xboole_0|k4_zfmisc_1(X43,X44,X45,X46)=k1_xboole_0)&((X45!=k1_xboole_0|k4_zfmisc_1(X43,X44,X45,X46)=k1_xboole_0)&(X46!=k1_xboole_0|k4_zfmisc_1(X43,X44,X45,X46)=k1_xboole_0))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t55_mcart_1])])])])).
cnf(c_0_42, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_44, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_45, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(esk9_0,k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0))), inference(rw,[status(thm)],[c_0_37, c_0_38])).
fof(c_0_47, plain, v1_xboole_0(esk10_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc1_xboole_0])])).
cnf(c_0_48, plain, (X1=k1_xboole_0|X2=k1_xboole_0|k2_zfmisc_1(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_49, plain, (X1=X2|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_50, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|k4_zfmisc_1(X1,X2,X3,X4)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_51, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_42, c_0_38])).
cnf(c_0_52, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_53, negated_conjecture, (r2_hidden(esk9_0,k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0))|v1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_54, plain, (v1_xboole_0(esk10_0)), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_55, plain, ![X72, X73]:~r2_hidden(X72,k2_zfmisc_1(X72,X73)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_zfmisc_1])])).
cnf(c_0_56, plain, (X1=k1_xboole_0|X2=k1_xboole_0|~v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_43])).
cnf(c_0_57, negated_conjecture, (esk7_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_58, plain, ![X58]:~v1_xboole_0(k1_zfmisc_1(X58)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
cnf(c_0_59, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_60, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0),X1)=k1_xboole_0|r2_hidden(esk9_0,k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0))), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_61, negated_conjecture, (esk6_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_62, negated_conjecture, (esk5_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_63, plain, (esk10_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_43, c_0_54])).
cnf(c_0_64, plain, (~r2_hidden(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_65, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_66, plain, ![X63, X64, X65]:((r2_hidden(k1_mcart_1(X63),X64)|~r2_hidden(X63,k2_zfmisc_1(X64,X65)))&(r2_hidden(k2_mcart_1(X63),X65)|~r2_hidden(X63,k2_zfmisc_1(X64,X65)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_mcart_1])])])).
cnf(c_0_67, negated_conjecture, (k2_zfmisc_1(esk5_0,esk6_0)=k1_xboole_0|r2_hidden(esk9_0,k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0))|~v1_xboole_0(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_53]), c_0_57])).
cnf(c_0_68, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_69, negated_conjecture, (X1=k1_xboole_0|r2_hidden(esk9_0,k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_57]), c_0_61]), c_0_62])).
cnf(c_0_70, plain, (v1_xboole_0(k1_xboole_0)), inference(rw,[status(thm)],[c_0_54, c_0_63])).
cnf(c_0_71, plain, (X1!=k1_xboole_0|~r2_hidden(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
fof(c_0_72, plain, ![X68, X69, X70, X71]:((k5_mcart_1(X68,X69,X70,X71)=k1_mcart_1(k1_mcart_1(X71))|~m1_subset_1(X71,k3_zfmisc_1(X68,X69,X70))|X70=k1_xboole_0|X69=k1_xboole_0|X68=k1_xboole_0)&((k6_mcart_1(X68,X69,X70,X71)=k2_mcart_1(k1_mcart_1(X71))|~m1_subset_1(X71,k3_zfmisc_1(X68,X69,X70))|X70=k1_xboole_0|X69=k1_xboole_0|X68=k1_xboole_0)&(k7_mcart_1(X68,X69,X70,X71)=k2_mcart_1(X71)|~m1_subset_1(X71,k3_zfmisc_1(X68,X69,X70))|X70=k1_xboole_0|X69=k1_xboole_0|X68=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_mcart_1])])])])])).
cnf(c_0_73, plain, (r2_hidden(k1_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_74, negated_conjecture, (k2_zfmisc_1(esk5_0,esk6_0)=k1_xboole_0|r2_hidden(esk9_0,k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0))), inference(spm,[status(thm)],[c_0_67, c_0_53])).
cnf(c_0_75, negated_conjecture, (r2_hidden(esk9_0,k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70])])).
cnf(c_0_76, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_77, plain, (~r2_hidden(X1,k1_xboole_0)), inference(er,[status(thm)],[c_0_71])).
cnf(c_0_78, plain, (k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_72])).
fof(c_0_79, plain, ![X74, X75, X76]:k3_mcart_1(X74,X75,X76)=k4_tarski(k4_tarski(X74,X75),X76), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_80, plain, ![X77, X78]:(~v1_relat_1(X78)|(~r2_hidden(X77,X78)|X77=k4_tarski(k1_mcart_1(X77),k2_mcart_1(X77)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_mcart_1])])).
cnf(c_0_81, negated_conjecture, (k2_zfmisc_1(esk5_0,esk6_0)=k1_xboole_0|r2_hidden(k1_mcart_1(esk9_0),k2_zfmisc_1(esk5_0,esk6_0))), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_82, negated_conjecture, (k2_zfmisc_1(esk5_0,esk6_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77])).
cnf(c_0_83, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_78, c_0_38])).
fof(c_0_84, plain, ![X66, X67]:v1_relat_1(k2_zfmisc_1(X66,X67)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_85, negated_conjecture, (esk8_0=X3|~m1_subset_1(X1,esk5_0)|~m1_subset_1(X2,esk6_0)|~m1_subset_1(X3,esk7_0)|esk9_0!=k3_mcart_1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_86, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_87, plain, (X2=k4_tarski(k1_mcart_1(X2),k2_mcart_1(X2))|~v1_relat_1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_88, negated_conjecture, (r2_hidden(k1_mcart_1(esk9_0),k2_zfmisc_1(esk5_0,esk6_0))), inference(sr,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_89, negated_conjecture, (k2_mcart_1(k1_mcart_1(esk9_0))=k6_mcart_1(esk5_0,esk6_0,esk7_0,esk9_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_46]), c_0_57]), c_0_61]), c_0_62])).
cnf(c_0_90, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_91, negated_conjecture, (esk8_0=X3|esk9_0!=k4_tarski(k4_tarski(X1,X2),X3)|~m1_subset_1(X3,esk7_0)|~m1_subset_1(X2,esk6_0)|~m1_subset_1(X1,esk5_0)), inference(rw,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_92, negated_conjecture, (k4_tarski(k1_mcart_1(k1_mcart_1(esk9_0)),k6_mcart_1(esk5_0,esk6_0,esk7_0,esk9_0))=k1_mcart_1(esk9_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_89]), c_0_90])])).
fof(c_0_93, plain, ![X79, X80]:(~r2_hidden(X79,X80)|m1_subset_1(X79,X80)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_94, negated_conjecture, (k2_zfmisc_1(esk5_0,esk6_0)=k1_xboole_0|r2_hidden(k1_mcart_1(k1_mcart_1(esk9_0)),esk5_0)), inference(spm,[status(thm)],[c_0_73, c_0_81])).
cnf(c_0_95, plain, (r2_hidden(k2_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_96, plain, (k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_97, negated_conjecture, (esk8_0=X1|k4_tarski(k1_mcart_1(esk9_0),X1)!=esk9_0|~m1_subset_1(k6_mcart_1(esk5_0,esk6_0,esk7_0,esk9_0),esk6_0)|~m1_subset_1(k1_mcart_1(k1_mcart_1(esk9_0)),esk5_0)|~m1_subset_1(X1,esk7_0)), inference(spm,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_98, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_99, negated_conjecture, (r2_hidden(k1_mcart_1(k1_mcart_1(esk9_0)),esk5_0)), inference(sr,[status(thm)],[c_0_94, c_0_82])).
cnf(c_0_100, negated_conjecture, (k2_zfmisc_1(esk5_0,esk6_0)=k1_xboole_0|r2_hidden(k6_mcart_1(esk5_0,esk6_0,esk7_0,esk9_0),esk6_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_81]), c_0_89])).
cnf(c_0_101, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_96, c_0_38])).
cnf(c_0_102, negated_conjecture, (esk8_0=X1|k4_tarski(k1_mcart_1(esk9_0),X1)!=esk9_0|~m1_subset_1(k6_mcart_1(esk5_0,esk6_0,esk7_0,esk9_0),esk6_0)|~m1_subset_1(X1,esk7_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_99])])).
cnf(c_0_103, negated_conjecture, (r2_hidden(k6_mcart_1(esk5_0,esk6_0,esk7_0,esk9_0),esk6_0)), inference(sr,[status(thm)],[c_0_100, c_0_82])).
cnf(c_0_104, negated_conjecture, (k4_tarski(k1_mcart_1(esk9_0),k2_mcart_1(esk9_0))=esk9_0|k2_zfmisc_1(esk5_0,esk6_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_74]), c_0_90])])).
cnf(c_0_105, negated_conjecture, (esk8_0!=k7_mcart_1(esk5_0,esk6_0,esk7_0,esk9_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_106, negated_conjecture, (k7_mcart_1(esk5_0,esk6_0,esk7_0,esk9_0)=k2_mcart_1(esk9_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_46]), c_0_57]), c_0_61]), c_0_62])).
cnf(c_0_107, negated_conjecture, (esk8_0=X1|k4_tarski(k1_mcart_1(esk9_0),X1)!=esk9_0|~m1_subset_1(X1,esk7_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_98]), c_0_103])])).
cnf(c_0_108, negated_conjecture, (k4_tarski(k1_mcart_1(esk9_0),k2_mcart_1(esk9_0))=esk9_0), inference(sr,[status(thm)],[c_0_104, c_0_82])).
cnf(c_0_109, negated_conjecture, (k2_mcart_1(esk9_0)!=esk8_0), inference(rw,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_110, negated_conjecture, (k2_zfmisc_1(esk5_0,esk6_0)=k1_xboole_0|r2_hidden(k2_mcart_1(esk9_0),esk7_0)), inference(spm,[status(thm)],[c_0_95, c_0_74])).
cnf(c_0_111, negated_conjecture, (~m1_subset_1(k2_mcart_1(esk9_0),esk7_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_109])).
cnf(c_0_112, negated_conjecture, (r2_hidden(k2_mcart_1(esk9_0),esk7_0)), inference(sr,[status(thm)],[c_0_110, c_0_82])).
cnf(c_0_113, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_98]), c_0_112])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 114
# Proof object clause steps            : 73
# Proof object formula steps           : 41
# Proof object conjectures             : 37
# Proof object clause conjectures      : 34
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 29
# Proof object simplifying inferences  : 40
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 46
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 43
# Processed clauses                    : 840
# ...of these trivial                  : 15
# ...subsumed                          : 572
# ...remaining for further processing  : 253
# Other redundant clauses eliminated   : 9
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 33
# Backward-rewritten                   : 10
# Generated clauses                    : 1935
# ...of the previous two non-trivial   : 1414
# Contextual simplify-reflections      : 1
# Paramodulations                      : 1909
# Factorizations                       : 0
# Equation resolutions                 : 17
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 160
#    Positive orientable unit clauses  : 19
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 17
#    Non-unit-clauses                  : 124
# Current number of unprocessed clauses: 473
# ...number of literals in the above   : 2092
# Current number of archived formulas  : 0
# Current number of archived clauses   : 96
# Clause-clause subsumption calls (NU) : 3014
# Rec. Clause-clause subsumption calls : 2278
# Non-unit clause-clause subsumptions  : 343
# Unit Clause-clause subsumption calls : 224
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 15
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 19748

# -------------------------------------------------
# User time                : 0.023 s
# System time              : 0.003 s
# Total time               : 0.026 s
# Maximum resident set size: 28780 pages
