# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B02_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(X2=k6_relat_1(X1)<=>![X3, X4]:(r2_hidden(k4_tarski(X3,X4),X2)<=>(r2_hidden(X3,X1)&X3=X4)))), file('data/MPTP/proofs/4bfe0aa9f6d140faa59e1b0217db3515.E_input', d10_relat_1)).
fof(d7_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(X2=k4_relat_1(X1)<=>![X3, X4]:(r2_hidden(k4_tarski(X3,X4),X2)<=>r2_hidden(k4_tarski(X4,X3),X1))))), file('data/MPTP/proofs/4bfe0aa9f6d140faa59e1b0217db3515.E_input', d7_relat_1)).
fof(dt_k6_relat_1, axiom, ![X1]:v1_relat_1(k6_relat_1(X1)), file('data/MPTP/proofs/4bfe0aa9f6d140faa59e1b0217db3515.E_input', dt_k6_relat_1)).
fof(t72_relat_1, conjecture, ![X1]:k4_relat_1(k6_relat_1(X1))=k6_relat_1(X1), file('data/MPTP/proofs/4bfe0aa9f6d140faa59e1b0217db3515.E_input', t72_relat_1)).
fof(c_0_4, plain, ![X15, X16, X17, X18, X19, X20]:((((r2_hidden(X17,X15)|~r2_hidden(k4_tarski(X17,X18),X16)|X16!=k6_relat_1(X15)|~v1_relat_1(X16))&(X17=X18|~r2_hidden(k4_tarski(X17,X18),X16)|X16!=k6_relat_1(X15)|~v1_relat_1(X16)))&(~r2_hidden(X19,X15)|X19!=X20|r2_hidden(k4_tarski(X19,X20),X16)|X16!=k6_relat_1(X15)|~v1_relat_1(X16)))&((~r2_hidden(k4_tarski(esk4_2(X15,X16),esk5_2(X15,X16)),X16)|(~r2_hidden(esk4_2(X15,X16),X15)|esk4_2(X15,X16)!=esk5_2(X15,X16))|X16=k6_relat_1(X15)|~v1_relat_1(X16))&((r2_hidden(esk4_2(X15,X16),X15)|r2_hidden(k4_tarski(esk4_2(X15,X16),esk5_2(X15,X16)),X16)|X16=k6_relat_1(X15)|~v1_relat_1(X16))&(esk4_2(X15,X16)=esk5_2(X15,X16)|r2_hidden(k4_tarski(esk4_2(X15,X16),esk5_2(X15,X16)),X16)|X16=k6_relat_1(X15)|~v1_relat_1(X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_relat_1])])])])])])).
fof(c_0_5, plain, ![X6, X7, X8, X9, X10, X11]:(((~r2_hidden(k4_tarski(X8,X9),X7)|r2_hidden(k4_tarski(X9,X8),X6)|X7!=k4_relat_1(X6)|~v1_relat_1(X7)|~v1_relat_1(X6))&(~r2_hidden(k4_tarski(X11,X10),X6)|r2_hidden(k4_tarski(X10,X11),X7)|X7!=k4_relat_1(X6)|~v1_relat_1(X7)|~v1_relat_1(X6)))&((~r2_hidden(k4_tarski(esk2_2(X6,X7),esk3_2(X6,X7)),X7)|~r2_hidden(k4_tarski(esk3_2(X6,X7),esk2_2(X6,X7)),X6)|X7=k4_relat_1(X6)|~v1_relat_1(X7)|~v1_relat_1(X6))&(r2_hidden(k4_tarski(esk2_2(X6,X7),esk3_2(X6,X7)),X7)|r2_hidden(k4_tarski(esk3_2(X6,X7),esk2_2(X6,X7)),X6)|X7=k4_relat_1(X6)|~v1_relat_1(X7)|~v1_relat_1(X6)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_relat_1])])])])])])).
cnf(c_0_6, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),X4)|X4!=k6_relat_1(X2)|~v1_relat_1(X4)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, plain, (r2_hidden(k4_tarski(esk2_2(X1,X2),esk3_2(X1,X2)),X2)|r2_hidden(k4_tarski(esk3_2(X1,X2),esk2_2(X1,X2)),X1)|X2=k4_relat_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_8, plain, ![X14]:v1_relat_1(k6_relat_1(X14)), inference(variable_rename,[status(thm)],[dt_k6_relat_1])).
cnf(c_0_9, plain, (X1=X2|~r2_hidden(k4_tarski(X1,X2),X3)|X3!=k6_relat_1(X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10, plain, (X1=k4_relat_1(X2)|r2_hidden(k4_tarski(esk2_2(X2,X1),esk3_2(X2,X1)),X1)|r2_hidden(esk3_2(X2,X1),X3)|X2!=k6_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_11, plain, (v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (esk3_2(X1,X2)=esk2_2(X1,X2)|X2=k4_relat_1(X1)|r2_hidden(k4_tarski(esk2_2(X1,X2),esk3_2(X1,X2)),X2)|X1!=k6_relat_1(X3)|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_9, c_0_7])).
cnf(c_0_13, plain, (X1=k4_relat_1(k6_relat_1(X2))|r2_hidden(k4_tarski(esk2_2(k6_relat_1(X2),X1),esk3_2(k6_relat_1(X2),X1)),X1)|r2_hidden(esk3_2(k6_relat_1(X2),X1),X2)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_10]), c_0_11])])).
cnf(c_0_14, plain, (esk3_2(k6_relat_1(X1),X2)=esk2_2(k6_relat_1(X1),X2)|X2=k4_relat_1(k6_relat_1(X1))|r2_hidden(k4_tarski(esk2_2(k6_relat_1(X1),X2),esk3_2(k6_relat_1(X1),X2)),X2)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_12]), c_0_11])])).
cnf(c_0_15, plain, (X1=k4_relat_1(k6_relat_1(X2))|r2_hidden(esk3_2(k6_relat_1(X2),X1),X2)|r2_hidden(esk2_2(k6_relat_1(X2),X1),X3)|X1!=k6_relat_1(X3)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_6, c_0_13])).
cnf(c_0_16, plain, (esk3_2(k6_relat_1(X1),X2)=esk2_2(k6_relat_1(X1),X2)|X2=k4_relat_1(k6_relat_1(X1))|X2!=k6_relat_1(X3)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_9, c_0_14])).
cnf(c_0_17, plain, (k6_relat_1(X1)=k4_relat_1(k6_relat_1(X2))|r2_hidden(esk2_2(k6_relat_1(X2),k6_relat_1(X1)),X1)|r2_hidden(esk3_2(k6_relat_1(X2),k6_relat_1(X1)),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_15]), c_0_11])])).
cnf(c_0_18, plain, (esk3_2(k6_relat_1(X1),k6_relat_1(X2))=esk2_2(k6_relat_1(X1),k6_relat_1(X2))|k6_relat_1(X2)=k4_relat_1(k6_relat_1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_16]), c_0_11])])).
cnf(c_0_19, plain, (r2_hidden(k4_tarski(X1,X3),X4)|~r2_hidden(X1,X2)|X1!=X3|X4!=k6_relat_1(X2)|~v1_relat_1(X4)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_20, plain, (k6_relat_1(X1)=k4_relat_1(k6_relat_1(X2))|r2_hidden(esk2_2(k6_relat_1(X2),k6_relat_1(X1)),X2)|r2_hidden(esk2_2(k6_relat_1(X2),k6_relat_1(X1)),X1)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_21, plain, (X2=k4_relat_1(X1)|~r2_hidden(k4_tarski(esk2_2(X1,X2),esk3_2(X1,X2)),X2)|~r2_hidden(k4_tarski(esk3_2(X1,X2),esk2_2(X1,X2)),X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_22, plain, (r2_hidden(k4_tarski(X1,X1),X2)|X2!=k6_relat_1(X3)|~r2_hidden(X1,X3)|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_23, plain, (k4_relat_1(k6_relat_1(X1))=k6_relat_1(X1)|r2_hidden(esk2_2(k6_relat_1(X1),k6_relat_1(X1)),X1)), inference(ef,[status(thm)],[c_0_20])).
fof(c_0_24, negated_conjecture, ~(![X1]:k4_relat_1(k6_relat_1(X1))=k6_relat_1(X1)), inference(assume_negation,[status(cth)],[t72_relat_1])).
cnf(c_0_25, plain, (k6_relat_1(X1)=k4_relat_1(k6_relat_1(X2))|~r2_hidden(k4_tarski(esk2_2(k6_relat_1(X2),k6_relat_1(X1)),esk2_2(k6_relat_1(X2),k6_relat_1(X1))),k6_relat_1(X1))|~r2_hidden(k4_tarski(esk2_2(k6_relat_1(X2),k6_relat_1(X1)),esk2_2(k6_relat_1(X2),k6_relat_1(X1))),k6_relat_1(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_18]), c_0_11]), c_0_11])])).
cnf(c_0_26, plain, (k4_relat_1(k6_relat_1(X1))=k6_relat_1(X1)|r2_hidden(k4_tarski(esk2_2(k6_relat_1(X1),k6_relat_1(X1)),esk2_2(k6_relat_1(X1),k6_relat_1(X1))),X2)|X2!=k6_relat_1(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
fof(c_0_27, negated_conjecture, k4_relat_1(k6_relat_1(esk1_0))!=k6_relat_1(esk1_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).
cnf(c_0_28, plain, (k4_relat_1(k6_relat_1(X1))=k6_relat_1(X1)|~r2_hidden(k4_tarski(esk2_2(k6_relat_1(X1),k6_relat_1(X1)),esk2_2(k6_relat_1(X1),k6_relat_1(X1))),k6_relat_1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_11])])).
cnf(c_0_29, negated_conjecture, (k4_relat_1(k6_relat_1(esk1_0))!=k6_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_30, plain, (k4_relat_1(k6_relat_1(X1))=k6_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_26]), c_0_11])])).
cnf(c_0_31, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_30])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 32
# Proof object clause steps            : 23
# Proof object formula steps           : 9
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 7
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 12
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 12
# Processed clauses                    : 73
# ...of these trivial                  : 0
# ...subsumed                          : 11
# ...remaining for further processing  : 62
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 15
# Generated clauses                    : 108
# ...of the previous two non-trivial   : 108
# Contextual simplify-reflections      : 0
# Paramodulations                      : 87
# Factorizations                       : 6
# Equation resolutions                 : 15
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 34
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 32
# Current number of unprocessed clauses: 48
# ...number of literals in the above   : 299
# Current number of archived formulas  : 0
# Current number of archived clauses   : 27
# Clause-clause subsumption calls (NU) : 305
# Rec. Clause-clause subsumption calls : 84
# Non-unit clause-clause subsumptions  : 11
# Unit Clause-clause subsumption calls : 22
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4254

# -------------------------------------------------
# User time                : 0.005 s
# System time              : 0.005 s
# Total time               : 0.011 s
# Maximum resident set size: 28632 pages
