# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t116_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>((v1_tsep_1(X2,X1)&m1_pre_topc(X2,X1))<=>g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=k8_tmap_1(X1,X2)))), file('data/MPTP/proofs/4cf1be8286d24317a2dbb433f137db43.E_input', t116_tmap_1)).
fof(t1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/4cf1be8286d24317a2dbb433f137db43.E_input', t1_tsep_1)).
fof(t114_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>(u1_struct_0(k8_tmap_1(X1,X2))=u1_struct_0(X1)&![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>u1_pre_topc(k8_tmap_1(X1,X2))=k5_tmap_1(X1,X3)))))), file('data/MPTP/proofs/4cf1be8286d24317a2dbb433f137db43.E_input', t114_tmap_1)).
fof(dt_k8_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_pre_topc(X2,X1))=>(v1_pre_topc(k8_tmap_1(X1,X2))&(v2_pre_topc(k8_tmap_1(X1,X2))&l1_pre_topc(k8_tmap_1(X1,X2))))), file('data/MPTP/proofs/4cf1be8286d24317a2dbb433f137db43.E_input', dt_k8_tmap_1)).
fof(t2_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>m1_pre_topc(X1,X1)), file('data/MPTP/proofs/4cf1be8286d24317a2dbb433f137db43.E_input', t2_tsep_1)).
fof(d5_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>r2_hidden(X2,u1_pre_topc(X1))))), file('data/MPTP/proofs/4cf1be8286d24317a2dbb433f137db43.E_input', d5_pre_topc)).
fof(t103_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X2,u1_pre_topc(X1))<=>u1_pre_topc(X1)=k5_tmap_1(X1,X2)))), file('data/MPTP/proofs/4cf1be8286d24317a2dbb433f137db43.E_input', t103_tmap_1)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/4cf1be8286d24317a2dbb433f137db43.E_input', dt_l1_pre_topc)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/4cf1be8286d24317a2dbb433f137db43.E_input', d3_struct_0)).
fof(fc9_tops_1, axiom, ![X1, X2]:(((v2_pre_topc(X1)&l1_pre_topc(X1))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>v3_pre_topc(k1_tops_1(X1,X2),X1)), file('data/MPTP/proofs/4cf1be8286d24317a2dbb433f137db43.E_input', fc9_tops_1)).
fof(abstractness_v1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v1_pre_topc(X1)=>X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))), file('data/MPTP/proofs/4cf1be8286d24317a2dbb433f137db43.E_input', abstractness_v1_pre_topc)).
fof(t43_tops_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>k1_tops_1(X1,k2_struct_0(X1))=k2_struct_0(X1)), file('data/MPTP/proofs/4cf1be8286d24317a2dbb433f137db43.E_input', t43_tops_1)).
fof(d1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>(v1_tsep_1(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>v3_pre_topc(X3,X1)))))), file('data/MPTP/proofs/4cf1be8286d24317a2dbb433f137db43.E_input', d1_tsep_1)).
fof(c_0_13, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>((v1_tsep_1(X2,X1)&m1_pre_topc(X2,X1))<=>g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=k8_tmap_1(X1,X2))))), inference(assume_negation,[status(cth)],[t116_tmap_1])).
fof(c_0_14, plain, ![X6, X7]:(~l1_pre_topc(X6)|(~m1_pre_topc(X7,X6)|m1_subset_1(u1_struct_0(X7),k1_zfmisc_1(u1_struct_0(X6))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])).
fof(c_0_15, plain, ![X8, X9, X10]:((u1_struct_0(k8_tmap_1(X8,X9))=u1_struct_0(X8)|(v2_struct_0(X9)|~m1_pre_topc(X9,X8))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))&(~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X8)))|(X10!=u1_struct_0(X9)|u1_pre_topc(k8_tmap_1(X8,X9))=k5_tmap_1(X8,X10))|(v2_struct_0(X9)|~m1_pre_topc(X9,X8))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t114_tmap_1])])])])])).
fof(c_0_16, plain, ![X11, X12]:((v1_pre_topc(k8_tmap_1(X11,X12))|(v2_struct_0(X11)|(~v2_pre_topc(X11)|~l1_pre_topc(X11))|~m1_pre_topc(X12,X11)))&((v2_pre_topc(k8_tmap_1(X11,X12))|(v2_struct_0(X11)|(~v2_pre_topc(X11)|~l1_pre_topc(X11))|~m1_pre_topc(X12,X11)))&(l1_pre_topc(k8_tmap_1(X11,X12))|(v2_struct_0(X11)|(~v2_pre_topc(X11)|~l1_pre_topc(X11))|~m1_pre_topc(X12,X11))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k8_tmap_1])])])])).
fof(c_0_17, plain, ![X13]:(~l1_pre_topc(X13)|m1_pre_topc(X13,X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tsep_1])])).
fof(c_0_18, plain, ![X14, X15]:((~v3_pre_topc(X15,X14)|r2_hidden(X15,u1_pre_topc(X14))|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X14)))|~l1_pre_topc(X14))&(~r2_hidden(X15,u1_pre_topc(X14))|v3_pre_topc(X15,X14)|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X14)))|~l1_pre_topc(X14))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_pre_topc])])])])).
fof(c_0_19, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((~v1_tsep_1(esk2_0,esk1_0)|~m1_pre_topc(esk2_0,esk1_0)|g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))!=k8_tmap_1(esk1_0,esk2_0))&((v1_tsep_1(esk2_0,esk1_0)|g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))=k8_tmap_1(esk1_0,esk2_0))&(m1_pre_topc(esk2_0,esk1_0)|g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))=k8_tmap_1(esk1_0,esk2_0)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_13])])])])])).
cnf(c_0_20, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (u1_struct_0(k8_tmap_1(X1,X2))=u1_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (l1_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (m1_pre_topc(X1,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_24, plain, ![X16, X17]:((~r2_hidden(X17,u1_pre_topc(X16))|u1_pre_topc(X16)=k5_tmap_1(X16,X17)|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))))&(u1_pre_topc(X16)!=k5_tmap_1(X16,X17)|r2_hidden(X17,u1_pre_topc(X16))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t103_tmap_1])])])])])).
cnf(c_0_25, plain, (r2_hidden(X1,u1_pre_topc(X2))|~v3_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(X1,k8_tmap_1(X2,X3))|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])).
cnf(c_0_28, plain, (m1_pre_topc(k8_tmap_1(X1,X2),k8_tmap_1(X1,X2))|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_23, c_0_22])).
cnf(c_0_29, plain, (u1_pre_topc(X2)=k5_tmap_1(X2,X1)|v2_struct_0(X2)|~r2_hidden(X1,u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, negated_conjecture, (r2_hidden(X1,u1_pre_topc(esk1_0))|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_31, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_32, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_33, plain, (m1_subset_1(u1_struct_0(k8_tmap_1(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
fof(c_0_34, plain, ![X18]:(~l1_pre_topc(X18)|l1_struct_0(X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_35, negated_conjecture, (u1_pre_topc(esk1_0)=k5_tmap_1(esk1_0,X1)|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_26]), c_0_31])]), c_0_32])).
fof(c_0_36, plain, ![X19]:(~l1_struct_0(X19)|k2_struct_0(X19)=u1_struct_0(X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
fof(c_0_37, plain, ![X20, X21]:(~v2_pre_topc(X20)|~l1_pre_topc(X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|v3_pre_topc(k1_tops_1(X20,X21),X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc9_tops_1])])).
fof(c_0_38, plain, ![X22]:(~l1_pre_topc(X22)|(~v1_pre_topc(X22)|X22=g1_pre_topc(u1_struct_0(X22),u1_pre_topc(X22)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_pre_topc])])).
cnf(c_0_39, plain, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_33, c_0_21])).
cnf(c_0_40, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_41, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_42, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43, negated_conjecture, (k5_tmap_1(esk1_0,u1_struct_0(X1))=u1_pre_topc(esk1_0)|~v3_pre_topc(u1_struct_0(X1),esk1_0)|~m1_pre_topc(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_20]), c_0_26])])).
cnf(c_0_44, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_45, plain, (v3_pre_topc(k1_tops_1(X1,X2),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_46, plain, ![X23]:(~v2_pre_topc(X23)|~l1_pre_topc(X23)|k1_tops_1(X23,k2_struct_0(X23))=k2_struct_0(X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_tops_1])])).
cnf(c_0_47, plain, (X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X1)|~v1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_48, plain, (v1_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_49, plain, (u1_pre_topc(k8_tmap_1(X2,X3))=k5_tmap_1(X2,X1)|v2_struct_0(X3)|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|X1!=u1_struct_0(X3)|~m1_pre_topc(X3,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_50, negated_conjecture, (m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_26]), c_0_31])]), c_0_32]), c_0_41])).
cnf(c_0_51, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))=k8_tmap_1(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_52, negated_conjecture, (l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_42, c_0_26])).
cnf(c_0_53, negated_conjecture, (k5_tmap_1(esk1_0,k2_struct_0(X1))=u1_pre_topc(esk1_0)|~l1_struct_0(X1)|~v3_pre_topc(k2_struct_0(X1),esk1_0)|~m1_pre_topc(X1,esk1_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_54, negated_conjecture, (m1_pre_topc(esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_23, c_0_26])).
cnf(c_0_55, plain, (v3_pre_topc(k1_tops_1(X1,u1_struct_0(X2)),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_45, c_0_20])).
cnf(c_0_56, plain, (k1_tops_1(X1,k2_struct_0(X1))=k2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_57, plain, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(k8_tmap_1(X1,X2)))=k8_tmap_1(X1,X2)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_21]), c_0_22]), c_0_48])).
cnf(c_0_58, negated_conjecture, (u1_pre_topc(k8_tmap_1(esk1_0,X1))=k5_tmap_1(esk1_0,u1_struct_0(esk1_0))|v2_struct_0(X1)|u1_struct_0(esk1_0)!=u1_struct_0(X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_26]), c_0_31])]), c_0_32])).
cnf(c_0_59, negated_conjecture, (g1_pre_topc(k2_struct_0(esk1_0),u1_pre_topc(esk1_0))=k8_tmap_1(esk1_0,esk2_0)|v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_44]), c_0_52])])).
cnf(c_0_60, negated_conjecture, (u1_pre_topc(esk1_0)=k5_tmap_1(esk1_0,k2_struct_0(esk1_0))|~v3_pre_topc(k2_struct_0(esk1_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_52])])).
cnf(c_0_61, plain, (v3_pre_topc(k1_tops_1(X1,k2_struct_0(X2)),X1)|~l1_struct_0(X2)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_55, c_0_44])).
cnf(c_0_62, negated_conjecture, (k1_tops_1(esk1_0,k2_struct_0(esk1_0))=k2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_31]), c_0_26])])).
cnf(c_0_63, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),k5_tmap_1(esk1_0,u1_struct_0(esk1_0)))=k8_tmap_1(esk1_0,X1)|v2_struct_0(X1)|u1_struct_0(esk1_0)!=u1_struct_0(X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_26]), c_0_31])]), c_0_32])).
fof(c_0_64, plain, ![X24, X25, X26]:((~v1_tsep_1(X25,X24)|(~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X24)))|(X26!=u1_struct_0(X25)|v3_pre_topc(X26,X24)))|~m1_pre_topc(X25,X24)|~l1_pre_topc(X24))&((m1_subset_1(esk3_2(X24,X25),k1_zfmisc_1(u1_struct_0(X24)))|v1_tsep_1(X25,X24)|~m1_pre_topc(X25,X24)|~l1_pre_topc(X24))&((esk3_2(X24,X25)=u1_struct_0(X25)|v1_tsep_1(X25,X24)|~m1_pre_topc(X25,X24)|~l1_pre_topc(X24))&(~v3_pre_topc(esk3_2(X24,X25),X24)|v1_tsep_1(X25,X24)|~m1_pre_topc(X25,X24)|~l1_pre_topc(X24))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tsep_1])])])])])).
cnf(c_0_65, negated_conjecture, (g1_pre_topc(k2_struct_0(esk1_0),k5_tmap_1(esk1_0,k2_struct_0(esk1_0)))=k8_tmap_1(esk1_0,esk2_0)|v1_tsep_1(esk2_0,esk1_0)|~v3_pre_topc(k2_struct_0(esk1_0),esk1_0)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_66, negated_conjecture, (v3_pre_topc(k2_struct_0(esk1_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_54]), c_0_62]), c_0_52]), c_0_26]), c_0_31])])).
cnf(c_0_67, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),k5_tmap_1(esk1_0,u1_struct_0(esk1_0)))=k8_tmap_1(esk1_0,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_54]), c_0_32])).
cnf(c_0_68, plain, (v3_pre_topc(X3,X2)|~v1_tsep_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|X3!=u1_struct_0(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_69, negated_conjecture, (g1_pre_topc(k2_struct_0(esk1_0),k5_tmap_1(esk1_0,k2_struct_0(esk1_0)))=k8_tmap_1(esk1_0,esk2_0)|v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_66])])).
cnf(c_0_70, negated_conjecture, (g1_pre_topc(k2_struct_0(esk1_0),k5_tmap_1(esk1_0,k2_struct_0(esk1_0)))=k8_tmap_1(esk1_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_44]), c_0_52])])).
cnf(c_0_71, plain, (r2_hidden(X2,u1_pre_topc(X1))|v2_struct_0(X1)|u1_pre_topc(X1)!=k5_tmap_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_72, plain, (v3_pre_topc(u1_struct_0(X1),X2)|u1_struct_0(X1)!=u1_struct_0(X3)|~v1_tsep_1(X3,X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_68, c_0_20])).
cnf(c_0_73, negated_conjecture, (k8_tmap_1(esk1_0,esk2_0)=k8_tmap_1(esk1_0,esk1_0)|v1_tsep_1(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_74, plain, (v3_pre_topc(X1,X2)|~r2_hidden(X1,u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_75, negated_conjecture, (r2_hidden(X1,u1_pre_topc(esk1_0))|u1_pre_topc(esk1_0)!=k5_tmap_1(esk1_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_31]), c_0_26])]), c_0_32])).
cnf(c_0_76, plain, (u1_pre_topc(k8_tmap_1(X1,X2))=k5_tmap_1(X1,u1_struct_0(X3))|v2_struct_0(X1)|v2_struct_0(X2)|u1_struct_0(X3)!=u1_struct_0(X2)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_49, c_0_20])).
cnf(c_0_77, negated_conjecture, (k8_tmap_1(esk1_0,esk2_0)=k8_tmap_1(esk1_0,esk1_0)|v3_pre_topc(u1_struct_0(X1),esk1_0)|u1_struct_0(X1)!=u1_struct_0(esk2_0)|~m1_pre_topc(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_40]), c_0_26])])).
cnf(c_0_78, negated_conjecture, (~v1_tsep_1(esk2_0,esk1_0)|~m1_pre_topc(esk2_0,esk1_0)|g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))!=k8_tmap_1(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_79, negated_conjecture, (v3_pre_topc(X1,esk1_0)|u1_pre_topc(esk1_0)!=k5_tmap_1(esk1_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_26])])).
cnf(c_0_80, plain, (u1_pre_topc(k8_tmap_1(X1,X2))=k5_tmap_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_76])).
cnf(c_0_81, negated_conjecture, (k8_tmap_1(esk1_0,esk2_0)=k8_tmap_1(esk1_0,esk1_0)|v3_pre_topc(u1_struct_0(esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_77, c_0_40])).
cnf(c_0_82, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))!=k8_tmap_1(esk1_0,esk2_0)|~v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_40])])).
cnf(c_0_83, negated_conjecture, (v3_pre_topc(u1_struct_0(X1),esk1_0)|k5_tmap_1(esk1_0,u1_struct_0(X1))!=u1_pre_topc(esk1_0)|~m1_pre_topc(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_20]), c_0_26])])).
cnf(c_0_84, negated_conjecture, (u1_pre_topc(esk1_0)=k5_tmap_1(esk1_0,k2_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_66])])).
cnf(c_0_85, negated_conjecture, (u1_pre_topc(k8_tmap_1(esk1_0,esk1_0))=k5_tmap_1(esk1_0,u1_struct_0(esk2_0))|v3_pre_topc(u1_struct_0(esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_40]), c_0_26]), c_0_31])]), c_0_41]), c_0_32])).
cnf(c_0_86, negated_conjecture, (g1_pre_topc(k2_struct_0(esk1_0),u1_pre_topc(esk1_0))!=k8_tmap_1(esk1_0,esk2_0)|~l1_struct_0(esk1_0)|~v1_tsep_1(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_82, c_0_44])).
cnf(c_0_87, negated_conjecture, (v3_pre_topc(u1_struct_0(X1),esk1_0)|k5_tmap_1(esk1_0,u1_struct_0(X1))!=k5_tmap_1(esk1_0,k2_struct_0(esk1_0))|~m1_pre_topc(X1,esk1_0)), inference(rw,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_88, negated_conjecture, (k5_tmap_1(esk1_0,u1_struct_0(esk2_0))=k5_tmap_1(esk1_0,u1_struct_0(esk1_0))|v3_pre_topc(u1_struct_0(esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_85]), c_0_54]), c_0_26]), c_0_31])]), c_0_32])).
cnf(c_0_89, negated_conjecture, (g1_pre_topc(k2_struct_0(esk1_0),u1_pre_topc(esk1_0))!=k8_tmap_1(esk1_0,esk2_0)|~v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_52])])).
cnf(c_0_90, plain, (g1_pre_topc(u1_struct_0(X1),k5_tmap_1(X1,u1_struct_0(X2)))=k8_tmap_1(X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_57, c_0_80])).
cnf(c_0_91, negated_conjecture, (k5_tmap_1(esk1_0,u1_struct_0(X1))=k5_tmap_1(esk1_0,k2_struct_0(esk1_0))|~v3_pre_topc(u1_struct_0(X1),esk1_0)|~m1_pre_topc(X1,esk1_0)), inference(rw,[status(thm)],[c_0_43, c_0_84])).
cnf(c_0_92, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),esk1_0)|k5_tmap_1(esk1_0,u1_struct_0(esk1_0))!=k5_tmap_1(esk1_0,k2_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_40])])).
cnf(c_0_93, negated_conjecture, (g1_pre_topc(k2_struct_0(esk1_0),k5_tmap_1(esk1_0,k2_struct_0(esk1_0)))!=k8_tmap_1(esk1_0,esk2_0)|~v1_tsep_1(esk2_0,esk1_0)), inference(rw,[status(thm)],[c_0_89, c_0_84])).
cnf(c_0_94, plain, (esk3_2(X1,X2)=u1_struct_0(X2)|v1_tsep_1(X2,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_95, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),k5_tmap_1(esk1_0,k2_struct_0(esk1_0)))=k8_tmap_1(esk1_0,X1)|v2_struct_0(X1)|~v3_pre_topc(u1_struct_0(X1),esk1_0)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_26]), c_0_31])]), c_0_32])).
cnf(c_0_96, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_44]), c_0_52])])).
cnf(c_0_97, negated_conjecture, (k8_tmap_1(esk1_0,esk2_0)!=k8_tmap_1(esk1_0,esk1_0)|~v1_tsep_1(esk2_0,esk1_0)), inference(rw,[status(thm)],[c_0_93, c_0_70])).
cnf(c_0_98, negated_conjecture, (u1_struct_0(X1)=esk3_2(esk1_0,X1)|v1_tsep_1(X1,esk1_0)|~m1_pre_topc(X1,esk1_0)), inference(spm,[status(thm)],[c_0_94, c_0_26])).
cnf(c_0_99, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),k5_tmap_1(esk1_0,k2_struct_0(esk1_0)))=k8_tmap_1(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_40])]), c_0_41])).
cnf(c_0_100, negated_conjecture, (u1_struct_0(esk2_0)=esk3_2(esk1_0,esk2_0)|k8_tmap_1(esk1_0,esk2_0)!=k8_tmap_1(esk1_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_40])])).
cnf(c_0_101, negated_conjecture, (k8_tmap_1(esk1_0,esk2_0)=k8_tmap_1(esk1_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_44]), c_0_70]), c_0_52])])).
cnf(c_0_102, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),k5_tmap_1(esk1_0,k2_struct_0(esk1_0)))!=k8_tmap_1(esk1_0,esk2_0)|~v3_pre_topc(k2_struct_0(esk1_0),esk1_0)|~v1_tsep_1(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_82, c_0_60])).
cnf(c_0_103, negated_conjecture, (u1_struct_0(esk2_0)=esk3_2(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_100, c_0_101])])).
cnf(c_0_104, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),k5_tmap_1(esk1_0,k2_struct_0(esk1_0)))!=k8_tmap_1(esk1_0,esk2_0)|~v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_102, c_0_66])])).
cnf(c_0_105, plain, (v1_tsep_1(X2,X1)|~v3_pre_topc(esk3_2(X1,X2),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_106, negated_conjecture, (v3_pre_topc(esk3_2(esk1_0,esk2_0),esk1_0)), inference(rw,[status(thm)],[c_0_96, c_0_103])).
cnf(c_0_107, negated_conjecture, (~v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_104, c_0_99])])).
cnf(c_0_108, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_40]), c_0_26])]), c_0_107]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 109
# Proof object clause steps            : 82
# Proof object formula steps           : 27
# Proof object conjectures             : 56
# Proof object clause conjectures      : 53
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 45
# Proof object simplifying inferences  : 96
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 28
# Processed clauses                    : 2285
# ...of these trivial                  : 48
# ...subsumed                          : 1481
# ...remaining for further processing  : 756
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 108
# Backward-rewritten                   : 302
# Generated clauses                    : 5719
# ...of the previous two non-trivial   : 5190
# Contextual simplify-reflections      : 128
# Paramodulations                      : 5706
# Factorizations                       : 0
# Equation resolutions                 : 11
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 317
#    Positive orientable unit clauses  : 24
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 289
# Current number of unprocessed clauses: 2633
# ...number of literals in the above   : 15710
# Current number of archived formulas  : 0
# Current number of archived clauses   : 439
# Clause-clause subsumption calls (NU) : 50972
# Rec. Clause-clause subsumption calls : 17760
# Non-unit clause-clause subsumptions  : 1380
# Unit Clause-clause subsumption calls : 1177
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 19
# BW rewrite match successes           : 17
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 140346

# -------------------------------------------------
# User time                : 0.106 s
# System time              : 0.000 s
# Total time               : 0.106 s
# Maximum resident set size: 28796 pages
