# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t159_funct_1, conjecture, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)<=>![X2]:?[X3]:r1_tarski(k10_relat_1(X1,k1_tarski(X2)),k1_tarski(X3)))), file('data/MPTP/proofs/4d35daa8532b40fe9aafaf0368a13c0d.E_input', t159_funct_1)).
fof(t144_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(![X2]:~((r2_hidden(X2,k2_relat_1(X1))&![X3]:~(k10_relat_1(X1,k1_tarski(X2))=k1_tarski(X3))))<=>v2_funct_1(X1))), file('data/MPTP/proofs/4d35daa8532b40fe9aafaf0368a13c0d.E_input', t144_funct_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/4d35daa8532b40fe9aafaf0368a13c0d.E_input', d3_tarski)).
fof(t142_funct_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r2_hidden(X1,k2_relat_1(X2))<=>~(k10_relat_1(X2,k1_tarski(X1))=k1_xboole_0))), file('data/MPTP/proofs/4d35daa8532b40fe9aafaf0368a13c0d.E_input', t142_funct_1)).
fof(t39_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(X1,k1_tarski(X2))<=>(X1=k1_xboole_0|X1=k1_tarski(X2))), file('data/MPTP/proofs/4d35daa8532b40fe9aafaf0368a13c0d.E_input', t39_zfmisc_1)).
fof(c_0_5, negated_conjecture, ~(![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)<=>![X2]:?[X3]:r1_tarski(k10_relat_1(X1,k1_tarski(X2)),k1_tarski(X3))))), inference(assume_negation,[status(cth)],[t159_funct_1])).
fof(c_0_6, negated_conjecture, ![X6, X7]:((v1_relat_1(esk1_0)&v1_funct_1(esk1_0))&((~v2_funct_1(esk1_0)|~r1_tarski(k10_relat_1(esk1_0,k1_tarski(esk2_0)),k1_tarski(X6)))&(v2_funct_1(esk1_0)|r1_tarski(k10_relat_1(esk1_0,k1_tarski(X7)),k1_tarski(esk3_1(X7)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
fof(c_0_7, plain, ![X11, X13, X14]:(((r2_hidden(esk4_1(X11),k2_relat_1(X11))|v2_funct_1(X11)|(~v1_relat_1(X11)|~v1_funct_1(X11)))&(k10_relat_1(X11,k1_tarski(esk4_1(X11)))!=k1_tarski(X13)|v2_funct_1(X11)|(~v1_relat_1(X11)|~v1_funct_1(X11))))&(~v2_funct_1(X11)|(~r2_hidden(X14,k2_relat_1(X11))|k10_relat_1(X11,k1_tarski(X14))=k1_tarski(esk5_2(X11,X14)))|(~v1_relat_1(X11)|~v1_funct_1(X11)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t144_funct_1])])])])])])).
fof(c_0_8, plain, ![X18, X19, X20, X21, X22]:((~r1_tarski(X18,X19)|(~r2_hidden(X20,X18)|r2_hidden(X20,X19)))&((r2_hidden(esk6_2(X21,X22),X21)|r1_tarski(X21,X22))&(~r2_hidden(esk6_2(X21,X22),X22)|r1_tarski(X21,X22)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_9, negated_conjecture, (~v2_funct_1(esk1_0)|~r1_tarski(k10_relat_1(esk1_0,k1_tarski(esk2_0)),k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (k10_relat_1(X1,k1_tarski(X2))=k1_tarski(esk5_2(X1,X2))|~v2_funct_1(X1)|~r2_hidden(X2,k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (r1_tarski(X1,X2)|~r2_hidden(esk6_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, negated_conjecture, (~r2_hidden(X1,k2_relat_1(X2))|~r1_tarski(k10_relat_1(esk1_0,k1_tarski(esk2_0)),k10_relat_1(X2,k1_tarski(X1)))|~v2_funct_1(esk1_0)|~v2_funct_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_14, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_15, negated_conjecture, (v1_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_16, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_17, plain, ![X16, X17]:((~r2_hidden(X16,k2_relat_1(X17))|k10_relat_1(X17,k1_tarski(X16))!=k1_xboole_0|~v1_relat_1(X17))&(k10_relat_1(X17,k1_tarski(X16))=k1_xboole_0|r2_hidden(X16,k2_relat_1(X17))|~v1_relat_1(X17))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t142_funct_1])])])])).
fof(c_0_18, plain, ![X9, X10]:((~r1_tarski(X9,k1_tarski(X10))|(X9=k1_xboole_0|X9=k1_tarski(X10)))&((X9!=k1_xboole_0|r1_tarski(X9,k1_tarski(X10)))&(X9!=k1_tarski(X10)|r1_tarski(X9,k1_tarski(X10))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t39_zfmisc_1])])])).
cnf(c_0_19, negated_conjecture, (~r2_hidden(esk2_0,k2_relat_1(esk1_0))|~v2_funct_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16])])).
cnf(c_0_20, plain, (k10_relat_1(X1,k1_tarski(X2))=k1_xboole_0|r2_hidden(X2,k2_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21, plain, (r1_tarski(X1,k1_tarski(X2))|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|~r1_tarski(X1,k1_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, negated_conjecture, (v2_funct_1(esk1_0)|r1_tarski(k10_relat_1(esk1_0,k1_tarski(X1)),k1_tarski(esk3_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_24, negated_conjecture, (k10_relat_1(esk1_0,k1_tarski(esk2_0))=k1_xboole_0|~v2_funct_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_16])])).
cnf(c_0_25, plain, (r1_tarski(k1_xboole_0,k1_tarski(X1))), inference(er,[status(thm)],[c_0_21])).
cnf(c_0_26, negated_conjecture, (k1_tarski(esk3_1(X1))=k10_relat_1(esk1_0,k1_tarski(X1))|k10_relat_1(esk1_0,k1_tarski(X1))=k1_xboole_0|v2_funct_1(esk1_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_27, negated_conjecture, (~v2_funct_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_24]), c_0_25])])).
cnf(c_0_28, plain, (v2_funct_1(X1)|k10_relat_1(X1,k1_tarski(esk4_1(X1)))!=k1_tarski(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_29, negated_conjecture, (k1_tarski(esk3_1(X1))=k10_relat_1(esk1_0,k1_tarski(X1))|k10_relat_1(esk1_0,k1_tarski(X1))=k1_xboole_0), inference(sr,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_30, plain, (~r2_hidden(X1,k2_relat_1(X2))|k10_relat_1(X2,k1_tarski(X1))!=k1_xboole_0|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_31, plain, (r2_hidden(esk4_1(X1),k2_relat_1(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_32, negated_conjecture, (k10_relat_1(esk1_0,k1_tarski(X1))=k1_xboole_0|v2_funct_1(X2)|k10_relat_1(X2,k1_tarski(esk4_1(X2)))!=k10_relat_1(esk1_0,k1_tarski(X1))|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_33, plain, (v2_funct_1(X1)|k10_relat_1(X1,k1_tarski(esk4_1(X1)))!=k1_xboole_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_34, negated_conjecture, (k10_relat_1(esk1_0,k1_tarski(esk4_1(esk1_0)))=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_32]), c_0_15]), c_0_16])]), c_0_27])).
cnf(c_0_35, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_15]), c_0_16])]), c_0_27]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 36
# Proof object clause steps            : 25
# Proof object formula steps           : 11
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 15
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 15
# Processed clauses                    : 67
# ...of these trivial                  : 1
# ...subsumed                          : 11
# ...remaining for further processing  : 55
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 113
# ...of the previous two non-trivial   : 94
# Contextual simplify-reflections      : 0
# Paramodulations                      : 107
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 34
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 27
# Current number of unprocessed clauses: 49
# ...number of literals in the above   : 243
# Current number of archived formulas  : 0
# Current number of archived clauses   : 19
# Clause-clause subsumption calls (NU) : 64
# Rec. Clause-clause subsumption calls : 31
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3302

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.003 s
# Total time               : 0.010 s
# Maximum resident set size: 28688 pages
