# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/4d46c080895949719e5a98c53dba2969.E_input', d2_tarski)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/4d46c080895949719e5a98c53dba2969.E_input', t70_enumset1)).
fof(t25_yellow_0, conjecture, ![X1]:((v3_orders_2(X1)&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X2=k12_lattice3(X1,X2,X3)<=>r3_orders_2(X1,X2,X3))))), file('data/MPTP/proofs/4d46c080895949719e5a98c53dba2969.E_input', t25_yellow_0)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/4d46c080895949719e5a98c53dba2969.E_input', redefinition_k6_domain_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/4d46c080895949719e5a98c53dba2969.E_input', t69_enumset1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/4d46c080895949719e5a98c53dba2969.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/4d46c080895949719e5a98c53dba2969.E_input', dt_l1_orders_2)).
fof(fc3_xboole_0, axiom, ![X1, X2]:~(v1_xboole_0(k2_tarski(X1,X2))), file('data/MPTP/proofs/4d46c080895949719e5a98c53dba2969.E_input', fc3_xboole_0)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/4d46c080895949719e5a98c53dba2969.E_input', redefinition_r3_orders_2)).
fof(t38_orders_2, axiom, ![X1]:((v3_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(~((?[X4]:((v6_orders_2(X4,X1)&m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1))))&(r2_hidden(X2,X4)&r2_hidden(X3,X4)))&(~(r1_orders_2(X1,X2,X3))&~(r1_orders_2(X1,X3,X2)))))&~(((r1_orders_2(X1,X2,X3)|r1_orders_2(X1,X3,X2))&![X4]:((v6_orders_2(X4,X1)&m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1))))=>~((r2_hidden(X2,X4)&r2_hidden(X3,X4)))))))))), file('data/MPTP/proofs/4d46c080895949719e5a98c53dba2969.E_input', t38_orders_2)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/4d46c080895949719e5a98c53dba2969.E_input', t4_subset)).
fof(t35_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v6_orders_2(k6_domain_1(u1_struct_0(X1),X2),X1)&m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/4d46c080895949719e5a98c53dba2969.E_input', t35_orders_2)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/4d46c080895949719e5a98c53dba2969.E_input', d1_xboole_0)).
fof(t23_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k12_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))))))), file('data/MPTP/proofs/4d46c080895949719e5a98c53dba2969.E_input', t23_yellow_0)).
fof(cc2_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v2_lattice3(X1)=>~(v2_struct_0(X1)))), file('data/MPTP/proofs/4d46c080895949719e5a98c53dba2969.E_input', cc2_lattice3)).
fof(c_0_15, plain, ![X9, X10, X11, X12, X13, X14, X15, X16]:(((~r2_hidden(X12,X11)|(X12=X9|X12=X10)|X11!=k2_tarski(X9,X10))&((X13!=X9|r2_hidden(X13,X11)|X11!=k2_tarski(X9,X10))&(X13!=X10|r2_hidden(X13,X11)|X11!=k2_tarski(X9,X10))))&(((esk4_3(X14,X15,X16)!=X14|~r2_hidden(esk4_3(X14,X15,X16),X16)|X16=k2_tarski(X14,X15))&(esk4_3(X14,X15,X16)!=X15|~r2_hidden(esk4_3(X14,X15,X16),X16)|X16=k2_tarski(X14,X15)))&(r2_hidden(esk4_3(X14,X15,X16),X16)|(esk4_3(X14,X15,X16)=X14|esk4_3(X14,X15,X16)=X15)|X16=k2_tarski(X14,X15)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_16, plain, ![X18, X19]:k1_enumset1(X18,X18,X19)=k2_tarski(X18,X19), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_17, negated_conjecture, ~(![X1]:((v3_orders_2(X1)&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X2=k12_lattice3(X1,X2,X3)<=>r3_orders_2(X1,X2,X3)))))), inference(assume_negation,[status(cth)],[t25_yellow_0])).
fof(c_0_18, plain, ![X43, X44]:(v1_xboole_0(X43)|~m1_subset_1(X44,X43)|k6_domain_1(X43,X44)=k1_tarski(X44)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_19, plain, ![X45]:k2_tarski(X45,X45)=k1_tarski(X45), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_20, plain, ![X46]:(v2_struct_0(X46)|~l1_struct_0(X46)|~v1_xboole_0(u1_struct_0(X46))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_21, plain, ![X47]:(~l1_orders_2(X47)|l1_struct_0(X47)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_22, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_24, plain, ![X31, X32]:~v1_xboole_0(k2_tarski(X31,X32)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc3_xboole_0])])).
fof(c_0_25, plain, ![X20, X21, X22]:((~r3_orders_2(X20,X21,X22)|r1_orders_2(X20,X21,X22)|(v2_struct_0(X20)|(~v3_orders_2(X20)|~l1_orders_2(X20))|(~m1_subset_1(X21,u1_struct_0(X20))|~m1_subset_1(X22,u1_struct_0(X20)))))&(~r1_orders_2(X20,X21,X22)|r3_orders_2(X20,X21,X22)|(v2_struct_0(X20)|(~v3_orders_2(X20)|~l1_orders_2(X20))|(~m1_subset_1(X21,u1_struct_0(X20))|~m1_subset_1(X22,u1_struct_0(X20)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
fof(c_0_26, negated_conjecture, ((v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v2_lattice3(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&((esk2_0!=k12_lattice3(esk1_0,esk2_0,esk3_0)|~r3_orders_2(esk1_0,esk2_0,esk3_0))&(esk2_0=k12_lattice3(esk1_0,esk2_0,esk3_0)|r3_orders_2(esk1_0,esk2_0,esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).
fof(c_0_27, plain, ![X23, X24, X25, X26]:((~v6_orders_2(X26,X23)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X23)))|(~r2_hidden(X24,X26)|~r2_hidden(X25,X26))|(r1_orders_2(X23,X24,X25)|r1_orders_2(X23,X25,X24))|~m1_subset_1(X25,u1_struct_0(X23))|~m1_subset_1(X24,u1_struct_0(X23))|(~v3_orders_2(X23)|~l1_orders_2(X23)))&((((v6_orders_2(esk5_3(X23,X24,X25),X23)|~r1_orders_2(X23,X24,X25)|~m1_subset_1(X25,u1_struct_0(X23))|~m1_subset_1(X24,u1_struct_0(X23))|(~v3_orders_2(X23)|~l1_orders_2(X23)))&(m1_subset_1(esk5_3(X23,X24,X25),k1_zfmisc_1(u1_struct_0(X23)))|~r1_orders_2(X23,X24,X25)|~m1_subset_1(X25,u1_struct_0(X23))|~m1_subset_1(X24,u1_struct_0(X23))|(~v3_orders_2(X23)|~l1_orders_2(X23))))&((r2_hidden(X24,esk5_3(X23,X24,X25))|~r1_orders_2(X23,X24,X25)|~m1_subset_1(X25,u1_struct_0(X23))|~m1_subset_1(X24,u1_struct_0(X23))|(~v3_orders_2(X23)|~l1_orders_2(X23)))&(r2_hidden(X25,esk5_3(X23,X24,X25))|~r1_orders_2(X23,X24,X25)|~m1_subset_1(X25,u1_struct_0(X23))|~m1_subset_1(X24,u1_struct_0(X23))|(~v3_orders_2(X23)|~l1_orders_2(X23)))))&(((v6_orders_2(esk5_3(X23,X24,X25),X23)|~r1_orders_2(X23,X25,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~m1_subset_1(X24,u1_struct_0(X23))|(~v3_orders_2(X23)|~l1_orders_2(X23)))&(m1_subset_1(esk5_3(X23,X24,X25),k1_zfmisc_1(u1_struct_0(X23)))|~r1_orders_2(X23,X25,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~m1_subset_1(X24,u1_struct_0(X23))|(~v3_orders_2(X23)|~l1_orders_2(X23))))&((r2_hidden(X24,esk5_3(X23,X24,X25))|~r1_orders_2(X23,X25,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~m1_subset_1(X24,u1_struct_0(X23))|(~v3_orders_2(X23)|~l1_orders_2(X23)))&(r2_hidden(X25,esk5_3(X23,X24,X25))|~r1_orders_2(X23,X25,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~m1_subset_1(X24,u1_struct_0(X23))|(~v3_orders_2(X23)|~l1_orders_2(X23))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t38_orders_2])])])])])])).
fof(c_0_28, plain, ![X28, X29, X30]:(~r2_hidden(X28,X29)|~m1_subset_1(X29,k1_zfmisc_1(X30))|m1_subset_1(X28,X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_29, plain, ![X48, X49]:((v6_orders_2(k6_domain_1(u1_struct_0(X48),X49),X48)|~m1_subset_1(X49,u1_struct_0(X48))|(v2_struct_0(X48)|(~v3_orders_2(X48)|~l1_orders_2(X48))))&(m1_subset_1(k6_domain_1(u1_struct_0(X48),X49),k1_zfmisc_1(u1_struct_0(X48)))|~m1_subset_1(X49,u1_struct_0(X48))|(v2_struct_0(X48)|(~v3_orders_2(X48)|~l1_orders_2(X48))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t35_orders_2])])])])])).
cnf(c_0_30, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_31, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_32, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_33, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_34, plain, (X1=X4|X1=X3|X2!=k1_enumset1(X3,X3,X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_22, c_0_23])).
fof(c_0_35, plain, ![X39, X40, X41]:((~v1_xboole_0(X39)|~r2_hidden(X40,X39))&(r2_hidden(esk7_1(X41),X41)|v1_xboole_0(X41))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_36, plain, (~v1_xboole_0(k2_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_37, plain, ![X33, X34, X35, X36, X37]:(((r1_orders_2(X33,X36,X34)|X36!=k12_lattice3(X33,X34,X35)|~m1_subset_1(X36,u1_struct_0(X33))|~m1_subset_1(X35,u1_struct_0(X33))|~m1_subset_1(X34,u1_struct_0(X33))|(~v5_orders_2(X33)|(~v2_lattice3(X33)|~l1_orders_2(X33))))&((r1_orders_2(X33,X36,X35)|X36!=k12_lattice3(X33,X34,X35)|~m1_subset_1(X36,u1_struct_0(X33))|~m1_subset_1(X35,u1_struct_0(X33))|~m1_subset_1(X34,u1_struct_0(X33))|(~v5_orders_2(X33)|(~v2_lattice3(X33)|~l1_orders_2(X33))))&(~m1_subset_1(X37,u1_struct_0(X33))|(~r1_orders_2(X33,X37,X34)|~r1_orders_2(X33,X37,X35)|r1_orders_2(X33,X37,X36))|X36!=k12_lattice3(X33,X34,X35)|~m1_subset_1(X36,u1_struct_0(X33))|~m1_subset_1(X35,u1_struct_0(X33))|~m1_subset_1(X34,u1_struct_0(X33))|(~v5_orders_2(X33)|(~v2_lattice3(X33)|~l1_orders_2(X33))))))&((m1_subset_1(esk6_4(X33,X34,X35,X36),u1_struct_0(X33))|~r1_orders_2(X33,X36,X35)|~r1_orders_2(X33,X36,X34)|X36=k12_lattice3(X33,X34,X35)|~m1_subset_1(X36,u1_struct_0(X33))|~m1_subset_1(X35,u1_struct_0(X33))|~m1_subset_1(X34,u1_struct_0(X33))|(~v5_orders_2(X33)|(~v2_lattice3(X33)|~l1_orders_2(X33))))&(((r1_orders_2(X33,esk6_4(X33,X34,X35,X36),X34)|~r1_orders_2(X33,X36,X35)|~r1_orders_2(X33,X36,X34)|X36=k12_lattice3(X33,X34,X35)|~m1_subset_1(X36,u1_struct_0(X33))|~m1_subset_1(X35,u1_struct_0(X33))|~m1_subset_1(X34,u1_struct_0(X33))|(~v5_orders_2(X33)|(~v2_lattice3(X33)|~l1_orders_2(X33))))&(r1_orders_2(X33,esk6_4(X33,X34,X35,X36),X35)|~r1_orders_2(X33,X36,X35)|~r1_orders_2(X33,X36,X34)|X36=k12_lattice3(X33,X34,X35)|~m1_subset_1(X36,u1_struct_0(X33))|~m1_subset_1(X35,u1_struct_0(X33))|~m1_subset_1(X34,u1_struct_0(X33))|(~v5_orders_2(X33)|(~v2_lattice3(X33)|~l1_orders_2(X33)))))&(~r1_orders_2(X33,esk6_4(X33,X34,X35,X36),X36)|~r1_orders_2(X33,X36,X35)|~r1_orders_2(X33,X36,X34)|X36=k12_lattice3(X33,X34,X35)|~m1_subset_1(X36,u1_struct_0(X33))|~m1_subset_1(X35,u1_struct_0(X33))|~m1_subset_1(X34,u1_struct_0(X33))|(~v5_orders_2(X33)|(~v2_lattice3(X33)|~l1_orders_2(X33))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_yellow_0])])])])])).
cnf(c_0_38, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_39, negated_conjecture, (esk2_0=k12_lattice3(esk1_0,esk2_0,esk3_0)|r3_orders_2(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_40, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_41, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_42, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_43, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_44, plain, (r1_orders_2(X2,X3,X4)|r1_orders_2(X2,X4,X3)|~v6_orders_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X3,X1)|~r2_hidden(X4,X1)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~v3_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_45, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_46, plain, (v6_orders_2(k6_domain_1(u1_struct_0(X1),X2),X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_47, plain, (k6_domain_1(X1,X2)=k1_enumset1(X2,X2,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_31]), c_0_23])).
cnf(c_0_48, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_49, plain, (m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_50, plain, (X1=X2|X1=X3|~r2_hidden(X1,k1_enumset1(X2,X2,X3))), inference(er,[status(thm)],[c_0_34])).
cnf(c_0_51, plain, (r2_hidden(esk7_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_52, plain, (~v1_xboole_0(k1_enumset1(X1,X1,X2))), inference(rw,[status(thm)],[c_0_36, c_0_23])).
cnf(c_0_53, plain, (X4=k12_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk6_4(X1,X2,X3,X4),X4)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_54, plain, (r1_orders_2(X1,esk6_4(X1,X2,X3,X4),X2)|X4=k12_lattice3(X1,X2,X3)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_55, plain, (r1_orders_2(X1,X2,X3)|X2!=k12_lattice3(X1,X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_56, negated_conjecture, (k12_lattice3(esk1_0,esk2_0,esk3_0)=esk2_0|r1_orders_2(esk1_0,esk2_0,esk3_0)|v2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40]), c_0_41]), c_0_42]), c_0_43])])).
cnf(c_0_57, negated_conjecture, (v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_58, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_59, plain, (r1_orders_2(X1,X2,X3)|r1_orders_2(X1,X3,X2)|~v6_orders_2(X4,X1)|~r2_hidden(X3,X4)|~r2_hidden(X2,X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_44, c_0_45]), c_0_45])).
cnf(c_0_60, plain, (v6_orders_2(k1_enumset1(X1,X1,X1),X2)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])).
cnf(c_0_61, plain, (v2_struct_0(X1)|m1_subset_1(k1_enumset1(X2,X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_47]), c_0_48])).
cnf(c_0_62, plain, (esk7_1(k1_enumset1(X1,X1,X2))=X2|esk7_1(k1_enumset1(X1,X1,X2))=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52])).
cnf(c_0_63, negated_conjecture, (esk2_0!=k12_lattice3(esk1_0,esk2_0,esk3_0)|~r3_orders_2(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_64, plain, (r3_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_65, plain, (k12_lattice3(X1,X2,X3)=X2|~r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_66, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)|r1_orders_2(esk1_0,X1,esk3_0)|v2_struct_0(esk1_0)|X1!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_41]), c_0_40]), c_0_42]), c_0_57]), c_0_58])])).
cnf(c_0_67, plain, (r1_orders_2(X1,X2,X3)|r1_orders_2(X1,X3,X2)|v2_struct_0(X1)|~r2_hidden(X3,k1_enumset1(X4,X4,X4))|~r2_hidden(X2,k1_enumset1(X4,X4,X4))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61])).
cnf(c_0_68, plain, (esk7_1(k1_enumset1(X1,X1,X1))=X1), inference(er,[status(thm)],[inference(ef,[status(thm)],[c_0_62])])).
fof(c_0_69, plain, ![X50]:(~l1_orders_2(X50)|(~v2_lattice3(X50)|~v2_struct_0(X50))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_lattice3])])])).
cnf(c_0_70, negated_conjecture, (v2_struct_0(esk1_0)|k12_lattice3(esk1_0,esk2_0,esk3_0)!=esk2_0|~r1_orders_2(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_40]), c_0_41]), c_0_42]), c_0_43])])).
cnf(c_0_71, negated_conjecture, (k12_lattice3(esk1_0,X1,X2)=X1|~r1_orders_2(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_58]), c_0_42]), c_0_57])])).
cnf(c_0_72, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)|v2_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_66, c_0_41])).
cnf(c_0_73, plain, (r1_orders_2(X1,X2,X3)|r1_orders_2(X1,X3,X2)|v2_struct_0(X1)|~r2_hidden(X3,k1_enumset1(X2,X2,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_51]), c_0_68]), c_0_68]), c_0_52])).
cnf(c_0_74, plain, (~l1_orders_2(X1)|~v2_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_75, negated_conjecture, (v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_41]), c_0_40])]), c_0_72])).
cnf(c_0_76, plain, (r1_orders_2(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_51]), c_0_68]), c_0_68])]), c_0_52])).
cnf(c_0_77, negated_conjecture, (~r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_42]), c_0_57])])).
cnf(c_0_78, negated_conjecture, (v2_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_41]), c_0_42]), c_0_43])]), c_0_77])).
cnf(c_0_79, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_78]), c_0_42]), c_0_57])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 80
# Proof object clause steps            : 49
# Proof object formula steps           : 31
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 51
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 44
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 42
# Processed clauses                    : 134
# ...of these trivial                  : 1
# ...subsumed                          : 5
# ...remaining for further processing  : 128
# Other redundant clauses eliminated   : 10
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 6
# Generated clauses                    : 114
# ...of the previous two non-trivial   : 75
# Contextual simplify-reflections      : 10
# Paramodulations                      : 96
# Factorizations                       : 2
# Equation resolutions                 : 16
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 75
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 63
# Current number of unprocessed clauses: 20
# ...number of literals in the above   : 152
# Current number of archived formulas  : 0
# Current number of archived clauses   : 53
# Clause-clause subsumption calls (NU) : 4021
# Rec. Clause-clause subsumption calls : 218
# Non-unit clause-clause subsumptions  : 15
# Unit Clause-clause subsumption calls : 58
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7219

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.000 s
# Total time               : 0.015 s
# Maximum resident set size: 28584 pages
