# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/4d521fde991a463f993835695b35622d.E_input', d4_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/4d521fde991a463f993835695b35622d.E_input', d3_tarski)).
fof(t134_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4)))), file('data/MPTP/proofs/4d521fde991a463f993835695b35622d.E_input', t134_zfmisc_1)).
fof(t122_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k3_xboole_0(X1,X2),X3)=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&k2_zfmisc_1(X3,k3_xboole_0(X1,X2))=k3_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))), file('data/MPTP/proofs/4d521fde991a463f993835695b35622d.E_input', t122_zfmisc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/4d521fde991a463f993835695b35622d.E_input', d10_xboole_0)).
fof(t123_zfmisc_1, axiom, ![X1, X2, X3, X4]:k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)), file('data/MPTP/proofs/4d521fde991a463f993835695b35622d.E_input', t123_zfmisc_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/4d521fde991a463f993835695b35622d.E_input', t28_xboole_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/4d521fde991a463f993835695b35622d.E_input', d5_xboole_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/4d521fde991a463f993835695b35622d.E_input', d1_xboole_0)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/4d521fde991a463f993835695b35622d.E_input', l13_xboole_0)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/4d521fde991a463f993835695b35622d.E_input', t113_zfmisc_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/4d521fde991a463f993835695b35622d.E_input', t3_xboole_0)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/4d521fde991a463f993835695b35622d.E_input', t65_xboole_1)).
fof(t7_tarski, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&r2_hidden(X4,X3))))))), file('data/MPTP/proofs/4d521fde991a463f993835695b35622d.E_input', t7_tarski)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/4d521fde991a463f993835695b35622d.E_input', t2_boole)).
fof(t6_xboole_0, axiom, ![X1, X2]:~((r2_xboole_0(X1,X2)&![X3]:~((r2_hidden(X3,X2)&~(r2_hidden(X3,X1)))))), file('data/MPTP/proofs/4d521fde991a463f993835695b35622d.E_input', t6_xboole_0)).
fof(d8_xboole_0, axiom, ![X1, X2]:(r2_xboole_0(X1,X2)<=>(r1_tarski(X1,X2)&~(X1=X2))), file('data/MPTP/proofs/4d521fde991a463f993835695b35622d.E_input', d8_xboole_0)).
fof(c_0_17, plain, ![X5, X6, X7, X8, X9, X10, X11, X12]:((((r2_hidden(X8,X5)|~r2_hidden(X8,X7)|X7!=k3_xboole_0(X5,X6))&(r2_hidden(X8,X6)|~r2_hidden(X8,X7)|X7!=k3_xboole_0(X5,X6)))&(~r2_hidden(X9,X5)|~r2_hidden(X9,X6)|r2_hidden(X9,X7)|X7!=k3_xboole_0(X5,X6)))&((~r2_hidden(esk1_3(X10,X11,X12),X12)|(~r2_hidden(esk1_3(X10,X11,X12),X10)|~r2_hidden(esk1_3(X10,X11,X12),X11))|X12=k3_xboole_0(X10,X11))&((r2_hidden(esk1_3(X10,X11,X12),X10)|r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k3_xboole_0(X10,X11))&(r2_hidden(esk1_3(X10,X11,X12),X11)|r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k3_xboole_0(X10,X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_18, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_19, plain, ![X14, X15, X16, X17, X18]:((~r1_tarski(X14,X15)|(~r2_hidden(X16,X14)|r2_hidden(X16,X15)))&((r2_hidden(esk2_2(X17,X18),X17)|r1_tarski(X17,X18))&(~r2_hidden(esk2_2(X17,X18),X18)|r1_tarski(X17,X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_20, negated_conjecture, ~(![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4))))), inference(assume_negation,[status(cth)],[t134_zfmisc_1])).
cnf(c_0_21, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_18])).
cnf(c_0_22, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_23, plain, ![X24, X25, X26]:(k2_zfmisc_1(k3_xboole_0(X24,X25),X26)=k3_xboole_0(k2_zfmisc_1(X24,X26),k2_zfmisc_1(X25,X26))&k2_zfmisc_1(X26,k3_xboole_0(X24,X25))=k3_xboole_0(k2_zfmisc_1(X26,X24),k2_zfmisc_1(X26,X25))), inference(variable_rename,[status(thm)],[t122_zfmisc_1])).
fof(c_0_24, negated_conjecture, (k2_zfmisc_1(esk3_0,esk4_0)=k2_zfmisc_1(esk5_0,esk6_0)&(esk3_0!=k1_xboole_0&(esk4_0!=k1_xboole_0&(esk3_0!=esk5_0|esk4_0!=esk6_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).
cnf(c_0_25, plain, (r1_tarski(X1,X2)|~r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (r1_tarski(k3_xboole_0(X1,X2),X3)|r2_hidden(esk2_2(k3_xboole_0(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27, plain, (k2_zfmisc_1(X1,k3_xboole_0(X2,X3))=k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28, negated_conjecture, (k2_zfmisc_1(esk3_0,esk4_0)=k2_zfmisc_1(esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_29, plain, ![X27, X28]:(((r1_tarski(X27,X28)|X27!=X28)&(r1_tarski(X28,X27)|X27!=X28))&(~r1_tarski(X27,X28)|~r1_tarski(X28,X27)|X27=X28)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_30, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_31, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk3_0,X1),k2_zfmisc_1(esk5_0,esk6_0))=k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk4_0))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_32, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),X3)=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_33, plain, ![X29, X30, X31, X32]:k2_zfmisc_1(k3_xboole_0(X29,X30),k3_xboole_0(X31,X32))=k3_xboole_0(k2_zfmisc_1(X29,X31),k2_zfmisc_1(X30,X32)), inference(variable_rename,[status(thm)],[t123_zfmisc_1])).
fof(c_0_34, plain, ![X33, X34]:(~r1_tarski(X33,X34)|k3_xboole_0(X33,X34)=X33), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_35, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, plain, (r1_tarski(k2_zfmisc_1(X1,k3_xboole_0(X2,X3)),k2_zfmisc_1(X1,X3))), inference(spm,[status(thm)],[c_0_30, c_0_27])).
cnf(c_0_37, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk3_0,esk5_0),esk6_0)=k2_zfmisc_1(esk3_0,k3_xboole_0(esk6_0,esk4_0))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_38, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_39, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_40, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_35])).
cnf(c_0_41, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk4_0)),k2_zfmisc_1(esk3_0,k3_xboole_0(esk6_0,esk4_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_31])).
cnf(c_0_42, plain, (k3_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_43, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk5_0,esk6_0),k2_zfmisc_1(X1,esk4_0))=k2_zfmisc_1(k3_xboole_0(esk3_0,X1),esk4_0)), inference(spm,[status(thm)],[c_0_32, c_0_28])).
cnf(c_0_44, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_45, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk5_0,esk6_0),k2_zfmisc_1(esk3_0,k3_xboole_0(esk6_0,esk4_0)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_28])).
cnf(c_0_46, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk4_0)),k2_zfmisc_1(esk5_0,esk6_0))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_47, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_48, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk3_0,esk5_0),esk4_0)=k2_zfmisc_1(esk5_0,k3_xboole_0(esk6_0,esk4_0))), inference(spm,[status(thm)],[c_0_27, c_0_43])).
cnf(c_0_49, negated_conjecture, (k2_zfmisc_1(esk3_0,k3_xboole_0(esk6_0,esk4_0))=k2_zfmisc_1(esk5_0,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])])).
cnf(c_0_50, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_47])).
fof(c_0_51, plain, ![X35, X36, X37, X38, X39, X40, X41, X42]:((((r2_hidden(X38,X35)|~r2_hidden(X38,X37)|X37!=k4_xboole_0(X35,X36))&(~r2_hidden(X38,X36)|~r2_hidden(X38,X37)|X37!=k4_xboole_0(X35,X36)))&(~r2_hidden(X39,X35)|r2_hidden(X39,X36)|r2_hidden(X39,X37)|X37!=k4_xboole_0(X35,X36)))&((~r2_hidden(esk7_3(X40,X41,X42),X42)|(~r2_hidden(esk7_3(X40,X41,X42),X40)|r2_hidden(esk7_3(X40,X41,X42),X41))|X42=k4_xboole_0(X40,X41))&((r2_hidden(esk7_3(X40,X41,X42),X40)|r2_hidden(esk7_3(X40,X41,X42),X42)|X42=k4_xboole_0(X40,X41))&(~r2_hidden(esk7_3(X40,X41,X42),X41)|r2_hidden(esk7_3(X40,X41,X42),X42)|X42=k4_xboole_0(X40,X41))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_52, plain, ![X44, X45, X46]:((~v1_xboole_0(X44)|~r2_hidden(X45,X44))&(r2_hidden(esk8_1(X46),X46)|v1_xboole_0(X46))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_53, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk5_0,k3_xboole_0(esk6_0,esk4_0)),k2_zfmisc_1(k3_xboole_0(esk3_0,esk5_0),X1))=k2_zfmisc_1(esk5_0,k3_xboole_0(esk6_0,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_48]), c_0_38]), c_0_28]), c_0_27])).
cnf(c_0_54, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk3_0,esk5_0),esk6_0)=k2_zfmisc_1(esk5_0,esk6_0)), inference(rw,[status(thm)],[c_0_37, c_0_49])).
cnf(c_0_55, plain, (r1_tarski(k3_xboole_0(X1,X2),X3)|r2_hidden(esk2_2(k3_xboole_0(X1,X2),X3),X1)), inference(spm,[status(thm)],[c_0_50, c_0_22])).
cnf(c_0_56, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_57, plain, (r2_hidden(esk8_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_58, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk5_0,k3_xboole_0(esk6_0,esk4_0)),k2_zfmisc_1(X1,esk4_0))=k2_zfmisc_1(k3_xboole_0(k3_xboole_0(esk3_0,esk5_0),X1),esk4_0)), inference(spm,[status(thm)],[c_0_32, c_0_48])).
cnf(c_0_59, negated_conjecture, (k2_zfmisc_1(esk5_0,k3_xboole_0(k3_xboole_0(esk6_0,esk4_0),esk6_0))=k2_zfmisc_1(esk5_0,esk6_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_27]), c_0_42])).
cnf(c_0_60, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_25, c_0_55])).
cnf(c_0_61, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_56])).
cnf(c_0_62, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|r2_hidden(esk8_1(k3_xboole_0(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_50, c_0_57])).
cnf(c_0_63, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|r2_hidden(esk8_1(k3_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_21, c_0_57])).
cnf(c_0_64, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(k3_xboole_0(esk3_0,esk5_0),esk3_0),esk4_0)=k2_zfmisc_1(esk5_0,esk6_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_28]), c_0_27]), c_0_59])).
cnf(c_0_65, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X1)=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_39, c_0_60])).
fof(c_0_66, plain, ![X48]:(~v1_xboole_0(X48)|X48=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_67, plain, (v1_xboole_0(k3_xboole_0(k4_xboole_0(X1,X2),X3))|~r2_hidden(esk8_1(k3_xboole_0(k4_xboole_0(X1,X2),X3)),X2)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_68, plain, (v1_xboole_0(k3_xboole_0(X1,k3_xboole_0(X2,X3)))|r2_hidden(esk8_1(k3_xboole_0(X1,k3_xboole_0(X2,X3))),X3)), inference(spm,[status(thm)],[c_0_21, c_0_63])).
fof(c_0_69, plain, ![X49, X50]:((k2_zfmisc_1(X49,X50)!=k1_xboole_0|(X49=k1_xboole_0|X50=k1_xboole_0))&((X49!=k1_xboole_0|k2_zfmisc_1(X49,X50)=k1_xboole_0)&(X50!=k1_xboole_0|k2_zfmisc_1(X49,X50)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_70, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_71, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(X1,esk4_0),k2_zfmisc_1(esk5_0,k3_xboole_0(esk6_0,esk4_0)))=k2_zfmisc_1(k3_xboole_0(X1,k3_xboole_0(esk3_0,esk5_0)),esk4_0)), inference(spm,[status(thm)],[c_0_32, c_0_48])).
cnf(c_0_72, negated_conjecture, (k2_zfmisc_1(esk5_0,k3_xboole_0(esk6_0,esk4_0))=k2_zfmisc_1(esk5_0,esk6_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_65]), c_0_48])).
cnf(c_0_73, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(X1,esk4_0),k2_zfmisc_1(esk5_0,esk6_0))=k2_zfmisc_1(k3_xboole_0(X1,esk3_0),esk4_0)), inference(spm,[status(thm)],[c_0_32, c_0_28])).
cnf(c_0_74, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_75, plain, (v1_xboole_0(k3_xboole_0(k4_xboole_0(X1,X2),k3_xboole_0(X3,X2)))), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_76, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_77, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_70])).
fof(c_0_78, plain, ![X51, X52, X54, X55, X56]:(((r2_hidden(esk9_2(X51,X52),X51)|r1_xboole_0(X51,X52))&(r2_hidden(esk9_2(X51,X52),X52)|r1_xboole_0(X51,X52)))&(~r2_hidden(X56,X54)|~r2_hidden(X56,X55)|~r1_xboole_0(X54,X55))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_79, plain, ![X57]:r1_xboole_0(X57,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_80, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(X1,k3_xboole_0(esk3_0,esk5_0)),esk4_0)=k2_zfmisc_1(k3_xboole_0(X1,esk3_0),esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_72]), c_0_73])).
cnf(c_0_81, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k1_xboole_0), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_82, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_76])).
cnf(c_0_83, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|r2_hidden(esk2_2(k4_xboole_0(X1,X2),X3),X1)), inference(spm,[status(thm)],[c_0_77, c_0_22])).
cnf(c_0_84, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_85, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_79])).
fof(c_0_86, plain, ![X58, X59, X61]:((r2_hidden(esk10_2(X58,X59),X59)|~r2_hidden(X58,X59))&(~r2_hidden(X61,X59)|~r2_hidden(X61,esk10_2(X58,X59))|~r2_hidden(X58,X59))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tarski])])])])])).
cnf(c_0_87, plain, (X1=k1_xboole_0|X2=k1_xboole_0|k2_zfmisc_1(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_88, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(k4_xboole_0(X1,esk5_0),esk3_0),esk4_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_82])).
cnf(c_0_89, negated_conjecture, (esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_90, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_25, c_0_83])).
cnf(c_0_91, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_92, plain, (r2_hidden(esk10_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_93, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_94, negated_conjecture, (k3_xboole_0(k4_xboole_0(X1,esk5_0),esk3_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_89])).
cnf(c_0_95, plain, (k3_xboole_0(k4_xboole_0(X1,X2),X1)=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_39, c_0_90])).
cnf(c_0_96, plain, (~r2_hidden(esk10_2(X1,k1_xboole_0),X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_97, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_93])).
cnf(c_0_98, negated_conjecture, (k4_xboole_0(esk3_0,esk5_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_99, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_96, c_0_92])).
cnf(c_0_100, negated_conjecture, (r2_hidden(X1,esk5_0)|~r2_hidden(X1,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_99])).
cnf(c_0_101, negated_conjecture, (r1_tarski(esk3_0,X1)|r2_hidden(esk2_2(esk3_0,X1),esk5_0)), inference(spm,[status(thm)],[c_0_100, c_0_22])).
cnf(c_0_102, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk5_0,esk3_0),esk4_0)=k2_zfmisc_1(esk5_0,k3_xboole_0(esk4_0,esk6_0))), inference(spm,[status(thm)],[c_0_27, c_0_73])).
cnf(c_0_103, negated_conjecture, (r1_tarski(esk3_0,esk5_0)), inference(spm,[status(thm)],[c_0_25, c_0_101])).
cnf(c_0_104, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk5_0,k3_xboole_0(X1,esk6_0)),k2_zfmisc_1(esk5_0,k3_xboole_0(esk4_0,esk6_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_102]), c_0_38]), c_0_28]), c_0_27])).
cnf(c_0_105, negated_conjecture, (k3_xboole_0(esk3_0,esk5_0)=esk3_0), inference(spm,[status(thm)],[c_0_39, c_0_103])).
cnf(c_0_106, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk5_0,esk6_0),k2_zfmisc_1(esk3_0,X1))=k2_zfmisc_1(esk3_0,k3_xboole_0(esk4_0,X1))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_107, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk5_0,esk6_0),k2_zfmisc_1(esk5_0,k3_xboole_0(esk4_0,esk6_0)))), inference(spm,[status(thm)],[c_0_104, c_0_42])).
cnf(c_0_108, negated_conjecture, (k2_zfmisc_1(esk3_0,esk6_0)=k2_zfmisc_1(esk5_0,esk6_0)), inference(rw,[status(thm)],[c_0_54, c_0_105])).
cnf(c_0_109, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk5_0,esk3_0),esk6_0)=k2_zfmisc_1(esk3_0,k3_xboole_0(esk4_0,esk6_0))), inference(spm,[status(thm)],[c_0_32, c_0_106])).
cnf(c_0_110, negated_conjecture, (k2_zfmisc_1(esk5_0,k3_xboole_0(esk4_0,esk6_0))=k2_zfmisc_1(esk5_0,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_107]), c_0_36])])).
cnf(c_0_111, negated_conjecture, (k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk4_0))=k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk6_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_108]), c_0_31])).
cnf(c_0_112, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_113, plain, (v1_xboole_0(k3_xboole_0(k4_xboole_0(X1,X2),X2))), inference(spm,[status(thm)],[c_0_67, c_0_63])).
cnf(c_0_114, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_115, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk3_0,k3_xboole_0(esk4_0,esk6_0)),k2_zfmisc_1(k3_xboole_0(esk5_0,esk3_0),X1))=k2_zfmisc_1(esk3_0,k3_xboole_0(esk4_0,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_109]), c_0_38]), c_0_106])).
cnf(c_0_116, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk5_0,esk3_0),esk4_0)=k2_zfmisc_1(esk5_0,esk6_0)), inference(rw,[status(thm)],[c_0_102, c_0_110])).
cnf(c_0_117, negated_conjecture, (k3_xboole_0(X1,esk4_0)=k1_xboole_0|k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk6_0))!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_111]), c_0_112])).
cnf(c_0_118, plain, (k3_xboole_0(k4_xboole_0(X1,X2),X2)=k1_xboole_0), inference(spm,[status(thm)],[c_0_74, c_0_113])).
cnf(c_0_119, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_114])).
cnf(c_0_120, negated_conjecture, (k2_zfmisc_1(esk3_0,k3_xboole_0(k3_xboole_0(esk4_0,esk6_0),esk4_0))=k2_zfmisc_1(esk5_0,esk6_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_31]), c_0_42]), c_0_28])).
cnf(c_0_121, negated_conjecture, (k3_xboole_0(k4_xboole_0(X1,esk6_0),esk4_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_119])])).
cnf(c_0_122, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(X1,esk6_0),k2_zfmisc_1(esk3_0,k3_xboole_0(esk4_0,esk6_0)))=k2_zfmisc_1(k3_xboole_0(X1,k3_xboole_0(esk5_0,esk3_0)),esk6_0)), inference(spm,[status(thm)],[c_0_32, c_0_109])).
cnf(c_0_123, negated_conjecture, (k2_zfmisc_1(esk3_0,k3_xboole_0(esk4_0,esk6_0))=k2_zfmisc_1(esk5_0,esk6_0)), inference(rw,[status(thm)],[c_0_120, c_0_65])).
cnf(c_0_124, negated_conjecture, (k4_xboole_0(esk4_0,esk6_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_95, c_0_121])).
cnf(c_0_125, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(X1,k3_xboole_0(esk5_0,esk3_0)),esk6_0)=k2_zfmisc_1(k3_xboole_0(X1,esk5_0),esk6_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_122, c_0_123]), c_0_32])).
cnf(c_0_126, negated_conjecture, (r2_hidden(X1,esk6_0)|~r2_hidden(X1,esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_124]), c_0_99])).
cnf(c_0_127, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(k4_xboole_0(X1,esk3_0),esk5_0),esk6_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_81]), c_0_82])).
fof(c_0_128, plain, ![X62]:k3_xboole_0(X62,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
cnf(c_0_129, negated_conjecture, (r1_tarski(esk4_0,X1)|r2_hidden(esk2_2(esk4_0,X1),esk6_0)), inference(spm,[status(thm)],[c_0_126, c_0_22])).
cnf(c_0_130, negated_conjecture, (k2_zfmisc_1(esk3_0,k3_xboole_0(k4_xboole_0(X1,esk4_0),esk6_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_118]), c_0_119])).
cnf(c_0_131, negated_conjecture, (k2_zfmisc_1(k4_xboole_0(esk5_0,esk3_0),esk6_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_127, c_0_95])).
cnf(c_0_132, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_128])).
cnf(c_0_133, negated_conjecture, (r1_tarski(esk4_0,esk6_0)), inference(spm,[status(thm)],[c_0_25, c_0_129])).
cnf(c_0_134, negated_conjecture, (k3_xboole_0(k4_xboole_0(X1,esk4_0),esk6_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_130]), c_0_112])).
cnf(c_0_135, negated_conjecture, (k2_zfmisc_1(k4_xboole_0(esk5_0,esk3_0),k3_xboole_0(X1,esk6_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_131]), c_0_132])).
cnf(c_0_136, negated_conjecture, (k3_xboole_0(esk4_0,esk6_0)=esk4_0), inference(spm,[status(thm)],[c_0_39, c_0_133])).
fof(c_0_137, plain, ![X63, X64]:((r2_hidden(esk11_2(X63,X64),X64)|~r2_xboole_0(X63,X64))&(~r2_hidden(esk11_2(X63,X64),X63)|~r2_xboole_0(X63,X64))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_xboole_0])])])])])).
cnf(c_0_138, negated_conjecture, (k4_xboole_0(esk6_0,esk4_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_95, c_0_134])).
cnf(c_0_139, negated_conjecture, (k2_zfmisc_1(k4_xboole_0(esk5_0,esk3_0),esk4_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_140, plain, (~r2_hidden(esk11_2(X1,X2),X1)|~r2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_137])).
cnf(c_0_141, negated_conjecture, (r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk6_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_138]), c_0_99])).
cnf(c_0_142, negated_conjecture, (k4_xboole_0(esk5_0,esk3_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_139]), c_0_89])).
cnf(c_0_143, negated_conjecture, (~r2_xboole_0(esk4_0,X1)|~r2_hidden(esk11_2(esk4_0,X1),esk6_0)), inference(spm,[status(thm)],[c_0_140, c_0_141])).
cnf(c_0_144, plain, (r2_hidden(esk11_2(X1,X2),X2)|~r2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_137])).
fof(c_0_145, plain, ![X66, X67]:(((r1_tarski(X66,X67)|~r2_xboole_0(X66,X67))&(X66!=X67|~r2_xboole_0(X66,X67)))&(~r1_tarski(X66,X67)|X66=X67|r2_xboole_0(X66,X67))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_xboole_0])])])])).
cnf(c_0_146, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk5_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_142]), c_0_99])).
cnf(c_0_147, negated_conjecture, (~r2_xboole_0(esk4_0,esk6_0)), inference(spm,[status(thm)],[c_0_143, c_0_144])).
cnf(c_0_148, plain, (X1=X2|r2_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_145])).
cnf(c_0_149, negated_conjecture, (~r2_xboole_0(esk3_0,X1)|~r2_hidden(esk11_2(esk3_0,X1),esk5_0)), inference(spm,[status(thm)],[c_0_140, c_0_146])).
cnf(c_0_150, negated_conjecture, (esk3_0!=esk5_0|esk4_0!=esk6_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_151, negated_conjecture, (esk4_0=esk6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_148]), c_0_133])])).
cnf(c_0_152, negated_conjecture, (~r2_xboole_0(esk3_0,esk5_0)), inference(spm,[status(thm)],[c_0_149, c_0_144])).
cnf(c_0_153, negated_conjecture, (esk3_0!=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_150, c_0_151])])).
cnf(c_0_154, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_148]), c_0_103])]), c_0_153]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 155
# Proof object clause steps            : 120
# Proof object formula steps           : 35
# Proof object conjectures             : 69
# Proof object clause conjectures      : 66
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 29
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 75
# Proof object simplifying inferences  : 62
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 44
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 44
# Processed clauses                    : 3253
# ...of these trivial                  : 81
# ...subsumed                          : 2511
# ...remaining for further processing  : 661
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 54
# Backward-rewritten                   : 264
# Generated clauses                    : 23667
# ...of the previous two non-trivial   : 13652
# Contextual simplify-reflections      : 0
# Paramodulations                      : 23644
# Factorizations                       : 9
# Equation resolutions                 : 12
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 330
#    Positive orientable unit clauses  : 110
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 18
#    Non-unit-clauses                  : 202
# Current number of unprocessed clauses: 10019
# ...number of literals in the above   : 21142
# Current number of archived formulas  : 0
# Current number of archived clauses   : 320
# Clause-clause subsumption calls (NU) : 9700
# Rec. Clause-clause subsumption calls : 8128
# Non-unit clause-clause subsumptions  : 554
# Unit Clause-clause subsumption calls : 4239
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 239
# BW rewrite match successes           : 82
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 237164

# -------------------------------------------------
# User time                : 0.121 s
# System time              : 0.003 s
# Total time               : 0.125 s
# Maximum resident set size: 28636 pages
