# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t53_funct_1, conjecture, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(?[X2]:((v1_relat_1(X2)&v1_funct_1(X2))&k5_relat_1(X1,X2)=k6_relat_1(k1_relat_1(X1)))=>v2_funct_1(X1))), file('data/MPTP/proofs/4db1c4e5d00e439e97c19971181338a1.E_input', t53_funct_1)).
fof(t71_relat_1, axiom, ![X1]:(k1_relat_1(k6_relat_1(X1))=X1&k2_relat_1(k6_relat_1(X1))=X1), file('data/MPTP/proofs/4db1c4e5d00e439e97c19971181338a1.E_input', t71_relat_1)).
fof(fc3_funct_1, axiom, ![X1]:(v1_relat_1(k6_relat_1(X1))&v1_funct_1(k6_relat_1(X1))), file('data/MPTP/proofs/4db1c4e5d00e439e97c19971181338a1.E_input', fc3_funct_1)).
fof(t46_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(k2_relat_1(X1),k1_relat_1(X2))=>k1_relat_1(k5_relat_1(X1,X2))=k1_relat_1(X1)))), file('data/MPTP/proofs/4db1c4e5d00e439e97c19971181338a1.E_input', t46_relat_1)).
fof(t55_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>k5_relat_1(k5_relat_1(X1,X2),X3)=k5_relat_1(X1,k5_relat_1(X2,X3))))), file('data/MPTP/proofs/4db1c4e5d00e439e97c19971181338a1.E_input', t55_relat_1)).
fof(t49_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)<=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((r1_tarski(k2_relat_1(X2),k1_relat_1(X1))&(r1_tarski(k2_relat_1(X3),k1_relat_1(X1))&(k1_relat_1(X2)=k1_relat_1(X3)&k5_relat_1(X2,X1)=k5_relat_1(X3,X1))))=>X2=X3))))), file('data/MPTP/proofs/4db1c4e5d00e439e97c19971181338a1.E_input', t49_funct_1)).
fof(t52_funct_1, axiom, ![X1]:v2_funct_1(k6_relat_1(X1)), file('data/MPTP/proofs/4db1c4e5d00e439e97c19971181338a1.E_input', t52_funct_1)).
fof(c_0_7, negated_conjecture, ~(![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(?[X2]:((v1_relat_1(X2)&v1_funct_1(X2))&k5_relat_1(X1,X2)=k6_relat_1(k1_relat_1(X1)))=>v2_funct_1(X1)))), inference(assume_negation,[status(cth)],[t53_funct_1])).
fof(c_0_8, plain, ![X6]:(k1_relat_1(k6_relat_1(X6))=X6&k2_relat_1(k6_relat_1(X6))=X6), inference(variable_rename,[status(thm)],[t71_relat_1])).
fof(c_0_9, negated_conjecture, ((v1_relat_1(esk1_0)&v1_funct_1(esk1_0))&(((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&k5_relat_1(esk1_0,esk2_0)=k6_relat_1(k1_relat_1(esk1_0)))&~v2_funct_1(esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
fof(c_0_10, plain, ![X7]:(v1_relat_1(k6_relat_1(X7))&v1_funct_1(k6_relat_1(X7))), inference(variable_rename,[status(thm)],[fc3_funct_1])).
fof(c_0_11, plain, ![X8, X9]:(~v1_relat_1(X8)|(~v1_relat_1(X9)|(~r1_tarski(k2_relat_1(X8),k1_relat_1(X9))|k1_relat_1(k5_relat_1(X8,X9))=k1_relat_1(X8)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t46_relat_1])])])).
cnf(c_0_12, plain, (k1_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, negated_conjecture, (k5_relat_1(esk1_0,esk2_0)=k6_relat_1(k1_relat_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, plain, ![X10, X11, X12]:(~v1_relat_1(X10)|(~v1_relat_1(X11)|(~v1_relat_1(X12)|k5_relat_1(k5_relat_1(X10,X11),X12)=k5_relat_1(X10,k5_relat_1(X11,X12))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_relat_1])])])).
fof(c_0_16, plain, ![X13, X14, X15]:((~v2_funct_1(X13)|(~v1_relat_1(X14)|~v1_funct_1(X14)|(~v1_relat_1(X15)|~v1_funct_1(X15)|(~r1_tarski(k2_relat_1(X14),k1_relat_1(X13))|(~r1_tarski(k2_relat_1(X15),k1_relat_1(X13))|(k1_relat_1(X14)!=k1_relat_1(X15)|k5_relat_1(X14,X13)!=k5_relat_1(X15,X13)))|X14=X15)))|(~v1_relat_1(X13)|~v1_funct_1(X13)))&(((v1_relat_1(esk3_1(X13))|v2_funct_1(X13)|(~v1_relat_1(X13)|~v1_funct_1(X13)))&(v1_funct_1(esk3_1(X13))|v2_funct_1(X13)|(~v1_relat_1(X13)|~v1_funct_1(X13))))&(((v1_relat_1(esk4_1(X13))|v2_funct_1(X13)|(~v1_relat_1(X13)|~v1_funct_1(X13)))&(v1_funct_1(esk4_1(X13))|v2_funct_1(X13)|(~v1_relat_1(X13)|~v1_funct_1(X13))))&(((r1_tarski(k2_relat_1(esk3_1(X13)),k1_relat_1(X13))|v2_funct_1(X13)|(~v1_relat_1(X13)|~v1_funct_1(X13)))&((r1_tarski(k2_relat_1(esk4_1(X13)),k1_relat_1(X13))|v2_funct_1(X13)|(~v1_relat_1(X13)|~v1_funct_1(X13)))&((k1_relat_1(esk3_1(X13))=k1_relat_1(esk4_1(X13))|v2_funct_1(X13)|(~v1_relat_1(X13)|~v1_funct_1(X13)))&(k5_relat_1(esk3_1(X13),X13)=k5_relat_1(esk4_1(X13),X13)|v2_funct_1(X13)|(~v1_relat_1(X13)|~v1_funct_1(X13))))))&(esk3_1(X13)!=esk4_1(X13)|v2_funct_1(X13)|(~v1_relat_1(X13)|~v1_funct_1(X13))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t49_funct_1])])])])])).
fof(c_0_17, plain, ![X18]:v2_funct_1(k6_relat_1(X18)), inference(variable_rename,[status(thm)],[t52_funct_1])).
cnf(c_0_18, plain, (k1_relat_1(k5_relat_1(X1,X2))=k1_relat_1(X1)|~v1_relat_1(X1)|~v1_relat_1(X2)|~r1_tarski(k2_relat_1(X1),k1_relat_1(X2))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19, negated_conjecture, (k1_relat_1(k5_relat_1(esk1_0,esk2_0))=k1_relat_1(esk1_0)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_20, negated_conjecture, (v1_relat_1(k5_relat_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_14, c_0_13])).
cnf(c_0_21, plain, (k5_relat_1(k5_relat_1(X1,X2),X3)=k5_relat_1(X1,k5_relat_1(X2,X3))|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (k5_relat_1(esk3_1(X1),X1)=k5_relat_1(esk4_1(X1),X1)|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (v1_relat_1(esk4_1(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (v2_funct_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (v1_funct_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_26, negated_conjecture, (k1_relat_1(k5_relat_1(X1,k5_relat_1(esk1_0,esk2_0)))=k1_relat_1(X1)|~r1_tarski(k2_relat_1(X1),k1_relat_1(esk1_0))|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])])).
cnf(c_0_27, plain, (r1_tarski(k2_relat_1(esk4_1(X1)),k1_relat_1(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28, negated_conjecture, (v1_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_29, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_30, negated_conjecture, (~v2_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_31, plain, (k5_relat_1(k5_relat_1(esk3_1(X1),X1),X2)=k5_relat_1(esk4_1(X1),k5_relat_1(X1,X2))|v2_funct_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])).
cnf(c_0_32, plain, (v1_relat_1(esk3_1(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_33, plain, (X2=X3|~v2_funct_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_funct_1(X3)|~r1_tarski(k2_relat_1(X2),k1_relat_1(X1))|~r1_tarski(k2_relat_1(X3),k1_relat_1(X1))|k1_relat_1(X2)!=k1_relat_1(X3)|k5_relat_1(X2,X1)!=k5_relat_1(X3,X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_34, negated_conjecture, (v2_funct_1(k5_relat_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_24, c_0_13])).
cnf(c_0_35, negated_conjecture, (v1_funct_1(k5_relat_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_25, c_0_13])).
cnf(c_0_36, plain, (r1_tarski(k2_relat_1(esk3_1(X1)),k1_relat_1(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_37, negated_conjecture, (k1_relat_1(k5_relat_1(esk4_1(esk1_0),k5_relat_1(esk1_0,esk2_0)))=k1_relat_1(esk4_1(esk1_0))|~v1_relat_1(esk4_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_29])]), c_0_30])).
cnf(c_0_38, plain, (k5_relat_1(esk4_1(X1),k5_relat_1(X1,X2))=k5_relat_1(esk3_1(X1),k5_relat_1(X1,X2))|v2_funct_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_31]), c_0_32])).
cnf(c_0_39, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_40, negated_conjecture, (X1=X2|k5_relat_1(X1,k5_relat_1(esk1_0,esk2_0))!=k5_relat_1(X2,k5_relat_1(esk1_0,esk2_0))|k1_relat_1(X1)!=k1_relat_1(X2)|~r1_tarski(k2_relat_1(X2),k1_relat_1(esk1_0))|~r1_tarski(k2_relat_1(X1),k1_relat_1(esk1_0))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_19]), c_0_34]), c_0_35]), c_0_20])])).
cnf(c_0_41, negated_conjecture, (k1_relat_1(k5_relat_1(esk3_1(esk1_0),k5_relat_1(esk1_0,esk2_0)))=k1_relat_1(esk3_1(esk1_0))|~v1_relat_1(esk3_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_36]), c_0_28]), c_0_29])]), c_0_30])).
cnf(c_0_42, negated_conjecture, (k1_relat_1(k5_relat_1(esk3_1(esk1_0),k5_relat_1(esk1_0,esk2_0)))=k1_relat_1(esk4_1(esk1_0))|~v1_relat_1(esk4_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_28]), c_0_39]), c_0_29])]), c_0_30])).
cnf(c_0_43, negated_conjecture, (X1=esk4_1(esk1_0)|k5_relat_1(X1,k5_relat_1(esk1_0,esk2_0))!=k5_relat_1(esk4_1(esk1_0),k5_relat_1(esk1_0,esk2_0))|k1_relat_1(X1)!=k1_relat_1(esk4_1(esk1_0))|~r1_tarski(k2_relat_1(X1),k1_relat_1(esk1_0))|~v1_funct_1(esk4_1(esk1_0))|~v1_funct_1(X1)|~v1_relat_1(esk4_1(esk1_0))|~v1_relat_1(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_27]), c_0_28]), c_0_29])]), c_0_30])).
cnf(c_0_44, negated_conjecture, (k1_relat_1(esk4_1(esk1_0))=k1_relat_1(esk3_1(esk1_0))|~v1_relat_1(esk3_1(esk1_0))|~v1_relat_1(esk4_1(esk1_0))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_45, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)|k5_relat_1(esk4_1(esk1_0),k5_relat_1(esk1_0,esk2_0))!=k5_relat_1(esk3_1(esk1_0),k5_relat_1(esk1_0,esk2_0))|~v1_funct_1(esk4_1(esk1_0))|~v1_funct_1(esk3_1(esk1_0))|~v1_relat_1(esk4_1(esk1_0))|~v1_relat_1(esk3_1(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_36]), c_0_28]), c_0_29])]), c_0_30]), c_0_44])).
cnf(c_0_46, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)|~v1_funct_1(esk4_1(esk1_0))|~v1_funct_1(esk3_1(esk1_0))|~v1_relat_1(esk4_1(esk1_0))|~v1_relat_1(esk3_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_38]), c_0_28]), c_0_39]), c_0_29])]), c_0_30])).
cnf(c_0_47, plain, (v1_funct_1(esk4_1(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_48, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)|~v1_funct_1(esk3_1(esk1_0))|~v1_relat_1(esk4_1(esk1_0))|~v1_relat_1(esk3_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_28]), c_0_29])]), c_0_30])).
cnf(c_0_49, plain, (v1_funct_1(esk3_1(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_50, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)|~v1_relat_1(esk4_1(esk1_0))|~v1_relat_1(esk3_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_28]), c_0_29])]), c_0_30])).
cnf(c_0_51, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)|~v1_relat_1(esk3_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_23]), c_0_28]), c_0_29])]), c_0_30])).
cnf(c_0_52, plain, (v2_funct_1(X1)|esk3_1(X1)!=esk4_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_53, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_32]), c_0_28]), c_0_29])]), c_0_30])).
cnf(c_0_54, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_28]), c_0_29])]), c_0_30]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 55
# Proof object clause steps            : 40
# Proof object formula steps           : 15
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 55
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 181
# ...of these trivial                  : 0
# ...subsumed                          : 32
# ...remaining for further processing  : 149
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 25
# Generated clauses                    : 770
# ...of the previous two non-trivial   : 726
# Contextual simplify-reflections      : 19
# Paramodulations                      : 763
# Factorizations                       : 0
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 97
#    Positive orientable unit clauses  : 16
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 80
# Current number of unprocessed clauses: 561
# ...number of literals in the above   : 5463
# Current number of archived formulas  : 0
# Current number of archived clauses   : 52
# Clause-clause subsumption calls (NU) : 1590
# Rec. Clause-clause subsumption calls : 236
# Non-unit clause-clause subsumptions  : 55
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 37598

# -------------------------------------------------
# User time                : 0.024 s
# System time              : 0.000 s
# Total time               : 0.024 s
# Maximum resident set size: 28632 pages
