# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(fraenkel_a_3_0_waybel_9, axiom, ![X1, X2, X3, X4]:(((~(v2_struct_0(X2))&l1_struct_0(X2))&((~(v2_struct_0(X3))&l1_waybel_0(X3,X2))&m1_subset_1(X4,u1_struct_0(X3))))=>(r2_hidden(X1,a_3_0_waybel_9(X2,X3,X4))<=>?[X5]:(m1_subset_1(X5,u1_struct_0(X3))&(X1=X5&r1_orders_2(X3,X4,X5))))), file('data/MPTP/proofs/4de8f49cca174bd2a4fb3e5b71bf9d77.E_input', fraenkel_a_3_0_waybel_9)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/4de8f49cca174bd2a4fb3e5b71bf9d77.E_input', d3_tarski)).
fof(t13_waybel_9, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>r1_tarski(u1_struct_0(k4_waybel_9(X1,X2,X3)),u1_struct_0(X2))))), file('data/MPTP/proofs/4de8f49cca174bd2a4fb3e5b71bf9d77.E_input', t13_waybel_9)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/4de8f49cca174bd2a4fb3e5b71bf9d77.E_input', t2_subset)).
fof(t12_waybel_9, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>u1_struct_0(k4_waybel_9(X1,X2,X3))=a_3_0_waybel_9(X1,X2,X3)))), file('data/MPTP/proofs/4de8f49cca174bd2a4fb3e5b71bf9d77.E_input', t12_waybel_9)).
fof(dt_l1_waybel_0, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(l1_waybel_0(X2,X1)=>l1_orders_2(X2))), file('data/MPTP/proofs/4de8f49cca174bd2a4fb3e5b71bf9d77.E_input', dt_l1_waybel_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/4de8f49cca174bd2a4fb3e5b71bf9d77.E_input', dt_l1_orders_2)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/4de8f49cca174bd2a4fb3e5b71bf9d77.E_input', fc2_struct_0)).
fof(c_0_8, plain, ![X24, X25, X26, X27, X29]:(((m1_subset_1(esk5_4(X24,X25,X26,X27),u1_struct_0(X26))|~r2_hidden(X24,a_3_0_waybel_9(X25,X26,X27))|(v2_struct_0(X25)|~l1_struct_0(X25)|(v2_struct_0(X26)|~l1_waybel_0(X26,X25)|~m1_subset_1(X27,u1_struct_0(X26)))))&((X24=esk5_4(X24,X25,X26,X27)|~r2_hidden(X24,a_3_0_waybel_9(X25,X26,X27))|(v2_struct_0(X25)|~l1_struct_0(X25)|(v2_struct_0(X26)|~l1_waybel_0(X26,X25)|~m1_subset_1(X27,u1_struct_0(X26)))))&(r1_orders_2(X26,X27,esk5_4(X24,X25,X26,X27))|~r2_hidden(X24,a_3_0_waybel_9(X25,X26,X27))|(v2_struct_0(X25)|~l1_struct_0(X25)|(v2_struct_0(X26)|~l1_waybel_0(X26,X25)|~m1_subset_1(X27,u1_struct_0(X26)))))))&(~m1_subset_1(X29,u1_struct_0(X26))|(X24!=X29|~r1_orders_2(X26,X27,X29))|r2_hidden(X24,a_3_0_waybel_9(X25,X26,X27))|(v2_struct_0(X25)|~l1_struct_0(X25)|(v2_struct_0(X26)|~l1_waybel_0(X26,X25)|~m1_subset_1(X27,u1_struct_0(X26)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_3_0_waybel_9])])])])])])).
fof(c_0_9, plain, ![X16, X17, X18, X19, X20]:((~r1_tarski(X16,X17)|(~r2_hidden(X18,X16)|r2_hidden(X18,X17)))&((r2_hidden(esk4_2(X19,X20),X19)|r1_tarski(X19,X20))&(~r2_hidden(esk4_2(X19,X20),X20)|r1_tarski(X19,X20)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_10, plain, (X1=esk5_4(X1,X2,X3,X4)|v2_struct_0(X2)|v2_struct_0(X3)|~r2_hidden(X1,a_3_0_waybel_9(X2,X3,X4))|~l1_struct_0(X2)|~l1_waybel_0(X3,X2)|~m1_subset_1(X4,u1_struct_0(X3))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_12, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>r1_tarski(u1_struct_0(k4_waybel_9(X1,X2,X3)),u1_struct_0(X2)))))), inference(assume_negation,[status(cth)],[t13_waybel_9])).
fof(c_0_13, plain, ![X13, X14]:(~m1_subset_1(X13,X14)|(v1_xboole_0(X14)|r2_hidden(X13,X14))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_14, plain, (m1_subset_1(esk5_4(X1,X2,X3,X4),u1_struct_0(X3))|v2_struct_0(X2)|v2_struct_0(X3)|~r2_hidden(X1,a_3_0_waybel_9(X2,X3,X4))|~l1_struct_0(X2)|~l1_waybel_0(X3,X2)|~m1_subset_1(X4,u1_struct_0(X3))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15, plain, (esk5_4(esk4_2(a_3_0_waybel_9(X1,X2,X3),X4),X1,X2,X3)=esk4_2(a_3_0_waybel_9(X1,X2,X3),X4)|r1_tarski(a_3_0_waybel_9(X1,X2,X3),X4)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
fof(c_0_16, negated_conjecture, ((~v2_struct_0(esk1_0)&l1_struct_0(esk1_0))&((~v2_struct_0(esk2_0)&l1_waybel_0(esk2_0,esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk2_0))&~r1_tarski(u1_struct_0(k4_waybel_9(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])).
fof(c_0_17, plain, ![X10, X11, X12]:(v2_struct_0(X10)|~l1_struct_0(X10)|(v2_struct_0(X11)|~l1_waybel_0(X11,X10)|(~m1_subset_1(X12,u1_struct_0(X11))|u1_struct_0(k4_waybel_9(X10,X11,X12))=a_3_0_waybel_9(X10,X11,X12)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t12_waybel_9])])])])).
cnf(c_0_18, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (r1_tarski(a_3_0_waybel_9(X1,X2,X3),X4)|m1_subset_1(esk4_2(a_3_0_waybel_9(X1,X2,X3),X4),u1_struct_0(X2))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_11])).
fof(c_0_20, plain, ![X22, X23]:(~l1_struct_0(X22)|(~l1_waybel_0(X23,X22)|l1_orders_2(X23))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_waybel_0])])])).
cnf(c_0_21, negated_conjecture, (~r1_tarski(u1_struct_0(k4_waybel_9(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (v2_struct_0(X1)|v2_struct_0(X2)|u1_struct_0(k4_waybel_9(X1,X2,X3))=a_3_0_waybel_9(X1,X2,X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, negated_conjecture, (l1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_27, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28, plain, (r1_tarski(X1,X2)|~r2_hidden(esk4_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_29, plain, (r2_hidden(esk4_2(a_3_0_waybel_9(X1,X2,X3),X4),u1_struct_0(X2))|v1_xboole_0(u1_struct_0(X2))|r1_tarski(a_3_0_waybel_9(X1,X2,X3),X4)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
fof(c_0_30, plain, ![X15]:(~l1_orders_2(X15)|l1_struct_0(X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_31, plain, (l1_orders_2(X2)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_32, plain, ![X9]:(v2_struct_0(X9)|~l1_struct_0(X9)|~v1_xboole_0(u1_struct_0(X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_33, negated_conjecture, (~r1_tarski(a_3_0_waybel_9(esk1_0,esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24]), c_0_25])]), c_0_26]), c_0_27])).
cnf(c_0_34, plain, (v1_xboole_0(u1_struct_0(X1))|r1_tarski(a_3_0_waybel_9(X2,X1,X3),u1_struct_0(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_36, negated_conjecture, (l1_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_24]), c_0_25])])).
cnf(c_0_37, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_23]), c_0_24]), c_0_25])]), c_0_27]), c_0_26])).
cnf(c_0_39, negated_conjecture, (l1_struct_0(esk2_0)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_40, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])]), c_0_26]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 41
# Proof object clause steps            : 24
# Proof object formula steps           : 17
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 18
# Processed clauses                    : 53
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 53
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 29
# ...of the previous two non-trivial   : 27
# Contextual simplify-reflections      : 1
# Paramodulations                      : 28
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 33
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 21
# Current number of unprocessed clauses: 10
# ...number of literals in the above   : 113
# Current number of archived formulas  : 0
# Current number of archived clauses   : 19
# Clause-clause subsumption calls (NU) : 612
# Rec. Clause-clause subsumption calls : 34
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 17
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2468

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.003 s
# Total time               : 0.010 s
# Maximum resident set size: 28788 pages
