# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t51_mcart_1, conjecture, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(~(X4=k5_mcart_1(X1,X2,X3,X4))&(~(X4=k6_mcart_1(X1,X2,X3,X4))&~(X4=k7_mcart_1(X1,X2,X3,X4)))))))))), file('data/MPTP/proofs/4e40abed39694f09bc3743e65a1c3854.E_input', t51_mcart_1)).
fof(t50_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))&(k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))&k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4))))))))), file('data/MPTP/proofs/4e40abed39694f09bc3743e65a1c3854.E_input', t50_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/4e40abed39694f09bc3743e65a1c3854.E_input', d3_zfmisc_1)).
fof(t48_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>X4=k3_mcart_1(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4),k7_mcart_1(X1,X2,X3,X4)))))))), file('data/MPTP/proofs/4e40abed39694f09bc3743e65a1c3854.E_input', t48_mcart_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/4e40abed39694f09bc3743e65a1c3854.E_input', d3_mcart_1)).
fof(t20_mcart_1, axiom, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1)))), file('data/MPTP/proofs/4e40abed39694f09bc3743e65a1c3854.E_input', t20_mcart_1)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/4e40abed39694f09bc3743e65a1c3854.E_input', t7_mcart_1)).
fof(t72_zfmisc_1, axiom, ![X1, X2, X3]:(k4_xboole_0(k2_tarski(X1,X2),X3)=k2_tarski(X1,X2)<=>(~(r2_hidden(X1,X3))&~(r2_hidden(X2,X3)))), file('data/MPTP/proofs/4e40abed39694f09bc3743e65a1c3854.E_input', t72_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/4e40abed39694f09bc3743e65a1c3854.E_input', t70_enumset1)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/4e40abed39694f09bc3743e65a1c3854.E_input', d2_tarski)).
fof(t65_zfmisc_1, axiom, ![X1, X2]:(k4_xboole_0(X1,k1_tarski(X2))=X1<=>~(r2_hidden(X2,X1))), file('data/MPTP/proofs/4e40abed39694f09bc3743e65a1c3854.E_input', t65_zfmisc_1)).
fof(t29_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4, X5]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k3_mcart_1(X3,X4,X5))))))), file('data/MPTP/proofs/4e40abed39694f09bc3743e65a1c3854.E_input', t29_mcart_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/4e40abed39694f09bc3743e65a1c3854.E_input', t113_zfmisc_1)).
fof(t10_mcart_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(r2_hidden(k1_mcart_1(X1),X2)&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/4e40abed39694f09bc3743e65a1c3854.E_input', t10_mcart_1)).
fof(c_0_14, negated_conjecture, ~(![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(~(X4=k5_mcart_1(X1,X2,X3,X4))&(~(X4=k6_mcart_1(X1,X2,X3,X4))&~(X4=k7_mcart_1(X1,X2,X3,X4))))))))))), inference(assume_negation,[status(cth)],[t51_mcart_1])).
fof(c_0_15, plain, ![X22, X23, X24, X25]:((k5_mcart_1(X22,X23,X24,X25)=k1_mcart_1(k1_mcart_1(X25))|~m1_subset_1(X25,k3_zfmisc_1(X22,X23,X24))|X24=k1_xboole_0|X23=k1_xboole_0|X22=k1_xboole_0)&((k6_mcart_1(X22,X23,X24,X25)=k2_mcart_1(k1_mcart_1(X25))|~m1_subset_1(X25,k3_zfmisc_1(X22,X23,X24))|X24=k1_xboole_0|X23=k1_xboole_0|X22=k1_xboole_0)&(k7_mcart_1(X22,X23,X24,X25)=k2_mcart_1(X25)|~m1_subset_1(X25,k3_zfmisc_1(X22,X23,X24))|X24=k1_xboole_0|X23=k1_xboole_0|X22=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_mcart_1])])])])])).
fof(c_0_16, plain, ![X45, X46, X47]:k3_zfmisc_1(X45,X46,X47)=k2_zfmisc_1(k2_zfmisc_1(X45,X46),X47), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_17, negated_conjecture, (esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&(m1_subset_1(esk4_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))&(esk4_0=k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)|(esk4_0=k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)|esk4_0=k7_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])).
fof(c_0_18, plain, ![X51, X52, X53, X54]:(X51=k1_xboole_0|(X52=k1_xboole_0|(X53=k1_xboole_0|(~m1_subset_1(X54,k3_zfmisc_1(X51,X52,X53))|X54=k3_mcart_1(k5_mcart_1(X51,X52,X53,X54),k6_mcart_1(X51,X52,X53,X54),k7_mcart_1(X51,X52,X53,X54)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t48_mcart_1])])])])).
fof(c_0_19, plain, ![X48, X49, X50]:k3_mcart_1(X48,X49,X50)=k4_tarski(k4_tarski(X48,X49),X50), inference(variable_rename,[status(thm)],[d3_mcart_1])).
cnf(c_0_20, plain, (k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, negated_conjecture, (m1_subset_1(esk4_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k3_mcart_1(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4),k7_mcart_1(X1,X2,X3,X4))|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_26, negated_conjecture, (m1_subset_1(esk4_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(rw,[status(thm)],[c_0_22, c_0_21])).
cnf(c_0_27, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_29, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_30, plain, ![X42, X43, X44]:((X42!=k1_mcart_1(X42)|X42!=k4_tarski(X43,X44))&(X42!=k2_mcart_1(X42)|X42!=k4_tarski(X43,X44))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t20_mcart_1])])])])])).
fof(c_0_31, plain, ![X26, X27]:(k1_mcart_1(k4_tarski(X26,X27))=X26&k2_mcart_1(k4_tarski(X26,X27))=X27), inference(variable_rename,[status(thm)],[t7_mcart_1])).
fof(c_0_32, plain, ![X39, X40, X41]:(((~r2_hidden(X39,X41)|k4_xboole_0(k2_tarski(X39,X40),X41)!=k2_tarski(X39,X40))&(~r2_hidden(X40,X41)|k4_xboole_0(k2_tarski(X39,X40),X41)!=k2_tarski(X39,X40)))&(r2_hidden(X39,X41)|r2_hidden(X40,X41)|k4_xboole_0(k2_tarski(X39,X40),X41)=k2_tarski(X39,X40))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t72_zfmisc_1])])])])).
fof(c_0_33, plain, ![X32, X33]:k1_enumset1(X32,X32,X33)=k2_tarski(X32,X33), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_34, plain, ![X13, X14, X15, X16, X17, X18, X19, X20]:(((~r2_hidden(X16,X15)|(X16=X13|X16=X14)|X15!=k2_tarski(X13,X14))&((X17!=X13|r2_hidden(X17,X15)|X15!=k2_tarski(X13,X14))&(X17!=X14|r2_hidden(X17,X15)|X15!=k2_tarski(X13,X14))))&(((esk5_3(X18,X19,X20)!=X18|~r2_hidden(esk5_3(X18,X19,X20),X20)|X20=k2_tarski(X18,X19))&(esk5_3(X18,X19,X20)!=X19|~r2_hidden(esk5_3(X18,X19,X20),X20)|X20=k2_tarski(X18,X19)))&(r2_hidden(esk5_3(X18,X19,X20),X20)|(esk5_3(X18,X19,X20)=X18|esk5_3(X18,X19,X20)=X19)|X20=k2_tarski(X18,X19)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_35, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X4=k4_tarski(k4_tarski(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4)),k7_mcart_1(X1,X2,X3,X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_21]), c_0_24])).
cnf(c_0_36, negated_conjecture, (k7_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=k2_mcart_1(esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_28]), c_0_29])).
cnf(c_0_37, plain, (X1!=k2_mcart_1(X1)|X1!=k4_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_39, plain, (~r2_hidden(X1,X2)|k4_xboole_0(k2_tarski(X3,X1),X2)!=k2_tarski(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_41, plain, ![X28, X29]:((k4_xboole_0(X28,k1_tarski(X29))!=X28|~r2_hidden(X29,X28))&(r2_hidden(X29,X28)|k4_xboole_0(X28,k1_tarski(X29))=X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t65_zfmisc_1])])])).
cnf(c_0_42, plain, (r2_hidden(X1,X2)|r2_hidden(X3,X2)|k4_xboole_0(k2_tarski(X1,X3),X2)=k2_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_44, plain, ![X34, X36, X37, X38]:((r2_hidden(esk6_1(X34),X34)|X34=k1_xboole_0)&((~r2_hidden(X36,X34)|esk6_1(X34)!=k3_mcart_1(X36,X37,X38)|X34=k1_xboole_0)&(~r2_hidden(X37,X34)|esk6_1(X34)!=k3_mcart_1(X36,X37,X38)|X34=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t29_mcart_1])])])])])])])).
cnf(c_0_45, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_46, negated_conjecture, (k4_tarski(k4_tarski(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)),k2_mcart_1(esk4_0))=esk4_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_26])]), c_0_27]), c_0_28]), c_0_29])).
cnf(c_0_47, negated_conjecture, (esk4_0=k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)|esk4_0=k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)|esk4_0=k7_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_48, plain, (k4_tarski(X1,X2)!=X2), inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_37]), c_0_38])).
cnf(c_0_49, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_50, plain, (k4_xboole_0(k1_enumset1(X3,X3,X1),X2)!=k1_enumset1(X3,X3,X1)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_40]), c_0_40])).
cnf(c_0_51, plain, (r2_hidden(X1,X2)|k4_xboole_0(X2,k1_tarski(X1))=X2), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_52, plain, (k4_xboole_0(X1,k1_tarski(X2))!=X1|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_53, plain, (k4_xboole_0(k1_enumset1(X1,X1,X3),X2)=k1_enumset1(X1,X1,X3)|r2_hidden(X3,X2)|r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_40]), c_0_40])).
cnf(c_0_54, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X4,X2)), inference(rw,[status(thm)],[c_0_43, c_0_40])).
cnf(c_0_55, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk6_1(X2)!=k3_mcart_1(X3,X1,X4)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_56, negated_conjecture, (k4_tarski(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0))=k1_mcart_1(esk4_0)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_57, negated_conjecture, (k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=esk4_0|k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=esk4_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_47]), c_0_26])]), c_0_48]), c_0_27]), c_0_28]), c_0_29])).
cnf(c_0_58, plain, (X1=X4|X1=X3|X2!=k1_enumset1(X3,X3,X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_49, c_0_40])).
cnf(c_0_59, plain, (r2_hidden(X1,k1_enumset1(X2,X2,X3))|~r2_hidden(X3,k1_tarski(X1))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_60, plain, (r2_hidden(esk6_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_61, plain, ![X30, X31]:((k2_zfmisc_1(X30,X31)!=k1_xboole_0|(X30=k1_xboole_0|X31=k1_xboole_0))&((X30!=k1_xboole_0|k2_zfmisc_1(X30,X31)=k1_xboole_0)&(X31!=k1_xboole_0|k2_zfmisc_1(X30,X31)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_62, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk6_1(X2)!=k3_mcart_1(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_63, plain, (r2_hidden(X1,k1_tarski(X2))|r2_hidden(X3,k1_tarski(X2))|~r2_hidden(X2,k1_enumset1(X1,X1,X3))), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_64, plain, (r2_hidden(X1,k1_enumset1(X2,X2,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_54])])).
cnf(c_0_65, plain, (X2=k1_xboole_0|esk6_1(X2)!=k4_tarski(k4_tarski(X3,X1),X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_55, c_0_24])).
cnf(c_0_66, negated_conjecture, (k4_tarski(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),esk4_0)=k1_mcart_1(esk4_0)|k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=esk4_0), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_67, plain, (X1=X2|X1=X3|~r2_hidden(X1,k1_enumset1(X3,X3,X2))), inference(er,[status(thm)],[c_0_58])).
cnf(c_0_68, plain, (k1_tarski(X1)=k1_xboole_0|r2_hidden(X1,k1_enumset1(X2,X2,esk6_1(k1_tarski(X1))))), inference(spm,[status(thm)],[c_0_59, c_0_60])).
fof(c_0_69, plain, ![X10, X11, X12]:((r2_hidden(k1_mcart_1(X10),X11)|~r2_hidden(X10,k2_zfmisc_1(X11,X12)))&(r2_hidden(k2_mcart_1(X10),X12)|~r2_hidden(X10,k2_zfmisc_1(X11,X12)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_mcart_1])])])).
cnf(c_0_70, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_71, plain, (X2=k1_xboole_0|esk6_1(X2)!=k4_tarski(k4_tarski(X1,X3),X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_62, c_0_24])).
cnf(c_0_72, negated_conjecture, (k4_tarski(k1_mcart_1(esk4_0),k2_mcart_1(esk4_0))=esk4_0), inference(rw,[status(thm)],[c_0_46, c_0_56])).
cnf(c_0_73, plain, (r2_hidden(X1,k1_tarski(X1))|r2_hidden(X2,k1_tarski(X1))), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_74, negated_conjecture, (k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=esk4_0|X1=k1_xboole_0|esk6_1(X1)!=k4_tarski(k1_mcart_1(esk4_0),X2)|~r2_hidden(esk4_0,X1)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_75, plain, (esk6_1(k1_tarski(X1))=X1|k1_tarski(X1)=k1_xboole_0|X1=X2), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_76, plain, (r2_hidden(k2_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_77, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_70])).
cnf(c_0_78, negated_conjecture, (X1=k1_xboole_0|esk6_1(X1)!=k4_tarski(esk4_0,X2)|~r2_hidden(k1_mcart_1(esk4_0),X1)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_79, plain, (r2_hidden(X1,k1_tarski(X1))), inference(ef,[status(thm)],[c_0_73])).
cnf(c_0_80, negated_conjecture, (k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=esk4_0|X1=k1_xboole_0|esk6_1(X1)!=esk4_0|~r2_hidden(esk4_0,X1)), inference(spm,[status(thm)],[c_0_74, c_0_72])).
cnf(c_0_81, plain, (esk6_1(k1_tarski(X1))=X1|k1_tarski(X1)=k1_xboole_0), inference(ef,[status(thm)],[c_0_75])).
cnf(c_0_82, negated_conjecture, (X1=k1_xboole_0|esk6_1(X1)!=k4_tarski(esk4_0,X2)|~r2_hidden(k2_mcart_1(esk4_0),X1)), inference(spm,[status(thm)],[c_0_65, c_0_46])).
cnf(c_0_83, plain, (r2_hidden(k2_mcart_1(X1),X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_84, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_85, plain, (k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_86, negated_conjecture, (k1_tarski(k1_mcart_1(esk4_0))=k1_xboole_0|esk6_1(k1_tarski(k1_mcart_1(esk4_0)))!=k4_tarski(esk4_0,X1)), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_87, negated_conjecture, (k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=esk4_0|k1_tarski(esk4_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81])]), c_0_79])])).
cnf(c_0_88, negated_conjecture, (X1=k1_xboole_0|esk6_1(X1)!=k4_tarski(esk4_0,X2)|~r2_hidden(esk4_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_89, plain, (r2_hidden(k1_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_90, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_84])).
cnf(c_0_91, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_85, c_0_21])).
cnf(c_0_92, negated_conjecture, (k1_tarski(k1_mcart_1(esk4_0))=k1_xboole_0|k1_mcart_1(esk4_0)!=k4_tarski(esk4_0,X1)), inference(spm,[status(thm)],[c_0_86, c_0_81])).
cnf(c_0_93, negated_conjecture, (k4_tarski(esk4_0,k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0))=k1_mcart_1(esk4_0)|k1_tarski(esk4_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_56, c_0_87])).
cnf(c_0_94, negated_conjecture, (k1_tarski(k4_tarski(esk4_0,X1))=k1_xboole_0|~r2_hidden(esk4_0,k1_xboole_0)), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_81])])).
cnf(c_0_95, plain, (r2_hidden(k1_mcart_1(X1),X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_96, negated_conjecture, (k1_mcart_1(k1_mcart_1(esk4_0))=k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_26]), c_0_27]), c_0_28]), c_0_29])).
cnf(c_0_97, negated_conjecture, (k1_tarski(k1_mcart_1(esk4_0))=k1_xboole_0|k1_tarski(esk4_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_98, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k1_xboole_0)), inference(spm,[status(thm)],[c_0_83, c_0_38])).
cnf(c_0_99, negated_conjecture, (r2_hidden(k4_tarski(esk4_0,X1),k1_xboole_0)|~r2_hidden(esk4_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_79, c_0_94])).
cnf(c_0_100, negated_conjecture, (r2_hidden(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),X1)|~r2_hidden(k1_mcart_1(esk4_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_101, negated_conjecture, (k1_tarski(esk4_0)=k1_xboole_0|r2_hidden(k1_mcart_1(esk4_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_79, c_0_97])).
cnf(c_0_102, negated_conjecture, (r2_hidden(X1,X2)|~r2_hidden(esk4_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_103, negated_conjecture, (k1_tarski(esk4_0)=k1_xboole_0|r2_hidden(esk4_0,X1)|~r2_hidden(k1_mcart_1(esk4_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_100, c_0_87])).
cnf(c_0_104, negated_conjecture, (r2_hidden(k1_mcart_1(esk4_0),k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_101]), c_0_102])).
cnf(c_0_105, negated_conjecture, (k1_tarski(esk4_0)=k1_xboole_0|r2_hidden(esk4_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_104])])).
cnf(c_0_106, negated_conjecture, (r2_hidden(esk4_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_105]), c_0_102])).
cnf(c_0_107, negated_conjecture, (r2_hidden(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_102, c_0_106])])).
cnf(c_0_108, plain, (X1=X2|X1=X3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_107])])).
cnf(c_0_109, plain, (X1=X2), inference(ef,[status(thm)],[c_0_108])).
cnf(c_0_110, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_29, c_0_109]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 111
# Proof object clause steps            : 82
# Proof object formula steps           : 29
# Proof object conjectures             : 37
# Proof object clause conjectures      : 34
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 35
# Proof object simplifying inferences  : 51
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 32
# Processed clauses                    : 575
# ...of these trivial                  : 2
# ...subsumed                          : 139
# ...remaining for further processing  : 434
# Other redundant clauses eliminated   : 147
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 51
# Backward-rewritten                   : 329
# Generated clauses                    : 7359
# ...of the previous two non-trivial   : 7158
# Contextual simplify-reflections      : 3
# Paramodulations                      : 7188
# Factorizations                       : 14
# Equation resolutions                 : 147
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 3
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 6431
# ...number of literals in the above   : 21405
# Current number of archived formulas  : 0
# Current number of archived clauses   : 427
# Clause-clause subsumption calls (NU) : 22177
# Rec. Clause-clause subsumption calls : 17046
# Non-unit clause-clause subsumptions  : 162
# Unit Clause-clause subsumption calls : 895
# Rewrite failures with RHS unbound    : 84
# BW rewrite match attempts            : 270
# BW rewrite match successes           : 264
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 104007

# -------------------------------------------------
# User time                : 0.069 s
# System time              : 0.003 s
# Total time               : 0.072 s
# Maximum resident set size: 28684 pages
