# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t16_mcart_1, conjecture, ![X1, X2, X3, X4]:(r2_hidden(X1,k2_zfmisc_1(X2,k2_tarski(X3,X4)))=>(r2_hidden(k1_mcart_1(X1),X2)&(k2_mcart_1(X1)=X3|k2_mcart_1(X1)=X4))), file('data/MPTP/proofs/4e46f771a1f04b3cbaf5311c26f2c7b2.E_input', t16_mcart_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/4e46f771a1f04b3cbaf5311c26f2c7b2.E_input', d5_xboole_0)).
fof(t144_zfmisc_1, axiom, ![X1, X2, X3]:~((~(r2_hidden(X1,X3))&(~(r2_hidden(X2,X3))&~(X3=k4_xboole_0(X3,k2_tarski(X1,X2)))))), file('data/MPTP/proofs/4e46f771a1f04b3cbaf5311c26f2c7b2.E_input', t144_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/4e46f771a1f04b3cbaf5311c26f2c7b2.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/4e46f771a1f04b3cbaf5311c26f2c7b2.E_input', t71_enumset1)).
fof(t10_mcart_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(r2_hidden(k1_mcart_1(X1),X2)&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/4e46f771a1f04b3cbaf5311c26f2c7b2.E_input', t10_mcart_1)).
fof(t38_zfmisc_1, axiom, ![X1, X2, X3]:(r1_tarski(k2_tarski(X1,X2),X3)<=>(r2_hidden(X1,X3)&r2_hidden(X2,X3))), file('data/MPTP/proofs/4e46f771a1f04b3cbaf5311c26f2c7b2.E_input', t38_zfmisc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/4e46f771a1f04b3cbaf5311c26f2c7b2.E_input', d10_xboole_0)).
fof(t12_mcart_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(k1_tarski(X2),X3))=>(k1_mcart_1(X1)=X2&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/4e46f771a1f04b3cbaf5311c26f2c7b2.E_input', t12_mcart_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/4e46f771a1f04b3cbaf5311c26f2c7b2.E_input', t69_enumset1)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/4e46f771a1f04b3cbaf5311c26f2c7b2.E_input', l54_zfmisc_1)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/4e46f771a1f04b3cbaf5311c26f2c7b2.E_input', t7_mcart_1)).
fof(c_0_12, negated_conjecture, ~(![X1, X2, X3, X4]:(r2_hidden(X1,k2_zfmisc_1(X2,k2_tarski(X3,X4)))=>(r2_hidden(k1_mcart_1(X1),X2)&(k2_mcart_1(X1)=X3|k2_mcart_1(X1)=X4)))), inference(assume_negation,[status(cth)],[t16_mcart_1])).
fof(c_0_13, plain, ![X31, X32, X33, X34, X35, X36, X37, X38]:((((r2_hidden(X34,X31)|~r2_hidden(X34,X33)|X33!=k4_xboole_0(X31,X32))&(~r2_hidden(X34,X32)|~r2_hidden(X34,X33)|X33!=k4_xboole_0(X31,X32)))&(~r2_hidden(X35,X31)|r2_hidden(X35,X32)|r2_hidden(X35,X33)|X33!=k4_xboole_0(X31,X32)))&((~r2_hidden(esk5_3(X36,X37,X38),X38)|(~r2_hidden(esk5_3(X36,X37,X38),X36)|r2_hidden(esk5_3(X36,X37,X38),X37))|X38=k4_xboole_0(X36,X37))&((r2_hidden(esk5_3(X36,X37,X38),X36)|r2_hidden(esk5_3(X36,X37,X38),X38)|X38=k4_xboole_0(X36,X37))&(~r2_hidden(esk5_3(X36,X37,X38),X37)|r2_hidden(esk5_3(X36,X37,X38),X38)|X38=k4_xboole_0(X36,X37))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_14, plain, ![X40, X41, X42]:(r2_hidden(X40,X42)|(r2_hidden(X41,X42)|X42=k4_xboole_0(X42,k2_tarski(X40,X41)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t144_zfmisc_1])])])).
fof(c_0_15, plain, ![X22, X23]:k1_enumset1(X22,X22,X23)=k2_tarski(X22,X23), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_16, plain, ![X16, X17, X18]:k2_enumset1(X16,X16,X17,X18)=k1_enumset1(X16,X17,X18), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_17, negated_conjecture, (r2_hidden(esk1_0,k2_zfmisc_1(esk2_0,k2_tarski(esk3_0,esk4_0)))&((k2_mcart_1(esk1_0)!=esk3_0|~r2_hidden(k1_mcart_1(esk1_0),esk2_0))&(k2_mcart_1(esk1_0)!=esk4_0|~r2_hidden(k1_mcart_1(esk1_0),esk2_0)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).
cnf(c_0_18, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (r2_hidden(X1,X2)|r2_hidden(X3,X2)|X2=k4_xboole_0(X2,k2_tarski(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, plain, ![X24, X25, X26]:((r2_hidden(k1_mcart_1(X24),X25)|~r2_hidden(X24,k2_zfmisc_1(X25,X26)))&(r2_hidden(k2_mcart_1(X24),X26)|~r2_hidden(X24,k2_zfmisc_1(X25,X26)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_mcart_1])])])).
cnf(c_0_23, negated_conjecture, (r2_hidden(esk1_0,k2_zfmisc_1(esk2_0,k2_tarski(esk3_0,esk4_0)))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (X2=k4_xboole_0(X2,k2_enumset1(X1,X1,X1,X3))|r2_hidden(X3,X2)|r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_20]), c_0_21])).
cnf(c_0_26, plain, (r2_hidden(k2_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27, negated_conjecture, (r2_hidden(esk1_0,k2_zfmisc_1(esk2_0,k2_enumset1(esk3_0,esk3_0,esk3_0,esk4_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_20]), c_0_21])).
fof(c_0_28, plain, ![X9, X10, X11]:(((r2_hidden(X9,X11)|~r1_tarski(k2_tarski(X9,X10),X11))&(r2_hidden(X10,X11)|~r1_tarski(k2_tarski(X9,X10),X11)))&(~r2_hidden(X9,X11)|~r2_hidden(X10,X11)|r1_tarski(k2_tarski(X9,X10),X11))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t38_zfmisc_1])])])).
fof(c_0_29, plain, ![X14, X15]:(((r1_tarski(X14,X15)|X14!=X15)&(r1_tarski(X15,X14)|X14!=X15))&(~r1_tarski(X14,X15)|~r1_tarski(X15,X14)|X14=X15)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_30, plain, (r2_hidden(X1,X2)|r2_hidden(X3,X2)|~r2_hidden(X4,k2_enumset1(X1,X1,X1,X3))|~r2_hidden(X4,X2)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31, negated_conjecture, (r2_hidden(k2_mcart_1(esk1_0),k2_enumset1(esk3_0,esk3_0,esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_tarski(X1,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34, negated_conjecture, (r2_hidden(esk4_0,X1)|r2_hidden(esk3_0,X1)|~r2_hidden(k2_mcart_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_35, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_enumset1(X1,X1,X1,X3),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_20]), c_0_21])).
cnf(c_0_36, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_33])).
cnf(c_0_37, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_38, negated_conjecture, (r2_hidden(esk3_0,k4_xboole_0(X1,X2))|~r2_hidden(k2_mcart_1(esk1_0),k4_xboole_0(X1,X2))|~r2_hidden(esk4_0,X2)), inference(spm,[status(thm)],[c_0_24, c_0_34])).
cnf(c_0_39, plain, (r2_hidden(X1,k2_enumset1(X1,X1,X1,X2))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
fof(c_0_40, plain, ![X19, X20, X21]:((k1_mcart_1(X19)=X20|~r2_hidden(X19,k2_zfmisc_1(k1_tarski(X20),X21)))&(r2_hidden(k2_mcart_1(X19),X21)|~r2_hidden(X19,k2_zfmisc_1(k1_tarski(X20),X21)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_mcart_1])])])).
fof(c_0_41, plain, ![X43]:k2_tarski(X43,X43)=k1_tarski(X43), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_42, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_43, negated_conjecture, (r2_hidden(esk3_0,k4_xboole_0(X1,k2_enumset1(esk4_0,esk4_0,esk4_0,X2)))|~r2_hidden(k2_mcart_1(esk1_0),k4_xboole_0(X1,k2_enumset1(esk4_0,esk4_0,esk4_0,X2)))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_44, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_45, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_tarski(X3,X1),X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_46, plain, (k1_mcart_1(X1)=X2|~r2_hidden(X1,k2_zfmisc_1(k1_tarski(X2),X3))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_47, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_48, plain, ![X27, X28, X29, X30]:(((r2_hidden(X27,X29)|~r2_hidden(k4_tarski(X27,X28),k2_zfmisc_1(X29,X30)))&(r2_hidden(X28,X30)|~r2_hidden(k4_tarski(X27,X28),k2_zfmisc_1(X29,X30))))&(~r2_hidden(X27,X29)|~r2_hidden(X28,X30)|r2_hidden(k4_tarski(X27,X28),k2_zfmisc_1(X29,X30)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
fof(c_0_49, plain, ![X12, X13]:(k1_mcart_1(k4_tarski(X12,X13))=X12&k2_mcart_1(k4_tarski(X12,X13))=X13), inference(variable_rename,[status(thm)],[t7_mcart_1])).
cnf(c_0_50, negated_conjecture, (r2_hidden(esk3_0,X1)|~r2_hidden(k2_mcart_1(esk1_0),k4_xboole_0(X1,k2_enumset1(esk4_0,esk4_0,esk4_0,X2)))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_51, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_52, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_enumset1(X3,X3,X3,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_20]), c_0_21])).
cnf(c_0_53, plain, (r2_hidden(k1_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_54, plain, (k1_mcart_1(X1)=X2|~r2_hidden(X1,k2_zfmisc_1(k2_enumset1(X2,X2,X2,X2),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_47]), c_0_20]), c_0_21])).
cnf(c_0_55, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_56, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_57, negated_conjecture, (r2_hidden(k2_mcart_1(esk1_0),k2_enumset1(esk4_0,esk4_0,esk4_0,X1))|r2_hidden(esk3_0,X2)|~r2_hidden(k2_mcart_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_58, plain, (r2_hidden(X1,k2_enumset1(X2,X2,X2,X1))), inference(spm,[status(thm)],[c_0_52, c_0_36])).
cnf(c_0_59, negated_conjecture, (k2_mcart_1(esk1_0)!=esk4_0|~r2_hidden(k1_mcart_1(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_60, negated_conjecture, (r2_hidden(k1_mcart_1(esk1_0),esk2_0)), inference(spm,[status(thm)],[c_0_53, c_0_27])).
cnf(c_0_61, plain, (X1=X2|~r2_hidden(X1,k2_enumset1(X2,X2,X2,X2))|~r2_hidden(X3,X4)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56])).
cnf(c_0_62, negated_conjecture, (r2_hidden(k2_mcart_1(esk1_0),k2_enumset1(esk4_0,esk4_0,esk4_0,X1))|r2_hidden(esk3_0,k2_enumset1(X2,X2,X2,k2_mcart_1(esk1_0)))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_63, negated_conjecture, (esk4_0!=k2_mcart_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_60])])).
cnf(c_0_64, negated_conjecture, (r2_hidden(esk3_0,k2_enumset1(X1,X1,X1,k2_mcart_1(esk1_0)))|~r2_hidden(X2,X3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63])).
cnf(c_0_65, negated_conjecture, (k2_mcart_1(esk1_0)!=esk3_0|~r2_hidden(k1_mcart_1(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_66, negated_conjecture, (r2_hidden(esk3_0,k2_enumset1(X1,X1,X1,k2_mcart_1(esk1_0)))), inference(spm,[status(thm)],[c_0_64, c_0_27])).
cnf(c_0_67, negated_conjecture, (esk3_0!=k2_mcart_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_60])])).
cnf(c_0_68, negated_conjecture, (r2_hidden(esk3_0,k2_enumset1(k2_mcart_1(esk1_0),k2_mcart_1(esk1_0),k2_mcart_1(esk1_0),X1))|r2_hidden(k2_mcart_1(esk1_0),k2_enumset1(esk4_0,esk4_0,esk4_0,X2))), inference(spm,[status(thm)],[c_0_57, c_0_39])).
cnf(c_0_69, negated_conjecture, (~r2_hidden(X1,X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_66]), c_0_67])).
cnf(c_0_70, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_68, c_0_69]), c_0_69]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 71
# Proof object clause steps            : 46
# Proof object formula steps           : 25
# Proof object conjectures             : 22
# Proof object clause conjectures      : 19
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 21
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 25
# Processed clauses                    : 99
# ...of these trivial                  : 0
# ...subsumed                          : 13
# ...remaining for further processing  : 86
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 4
# Generated clauses                    : 206
# ...of the previous two non-trivial   : 152
# Contextual simplify-reflections      : 0
# Paramodulations                      : 184
# Factorizations                       : 2
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 43
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 37
# Current number of unprocessed clauses: 96
# ...number of literals in the above   : 277
# Current number of archived formulas  : 0
# Current number of archived clauses   : 44
# Clause-clause subsumption calls (NU) : 533
# Rec. Clause-clause subsumption calls : 396
# Non-unit clause-clause subsumptions  : 14
# Unit Clause-clause subsumption calls : 54
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 10
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4742

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28608 pages
