# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PI_S0a
# and selection function SelectNegativeLiterals.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/4ebb96184ee44de792eb5f5fc9713dc7.E_input', d3_tarski)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/4ebb96184ee44de792eb5f5fc9713dc7.E_input', t28_xboole_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/4ebb96184ee44de792eb5f5fc9713dc7.E_input', d7_xboole_0)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/4ebb96184ee44de792eb5f5fc9713dc7.E_input', t3_xboole_0)).
fof(t38_subset_1, conjecture, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(r1_tarski(X2,k3_subset_1(X1,X2))<=>X2=k1_subset_1(X1))), file('data/MPTP/proofs/4ebb96184ee44de792eb5f5fc9713dc7.E_input', t38_subset_1)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/4ebb96184ee44de792eb5f5fc9713dc7.E_input', d1_zfmisc_1)).
fof(d3_subset_1, axiom, ![X1]:k1_subset_1(X1)=k1_xboole_0, file('data/MPTP/proofs/4ebb96184ee44de792eb5f5fc9713dc7.E_input', d3_subset_1)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/4ebb96184ee44de792eb5f5fc9713dc7.E_input', d4_xboole_0)).
fof(dt_k3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>m1_subset_1(k3_subset_1(X1,X2),k1_zfmisc_1(X1))), file('data/MPTP/proofs/4ebb96184ee44de792eb5f5fc9713dc7.E_input', dt_k3_subset_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/4ebb96184ee44de792eb5f5fc9713dc7.E_input', d2_subset_1)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/4ebb96184ee44de792eb5f5fc9713dc7.E_input', fc1_subset_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/4ebb96184ee44de792eb5f5fc9713dc7.E_input', t7_xboole_0)).
fof(t85_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_xboole_0(X1,k4_xboole_0(X3,X2))), file('data/MPTP/proofs/4ebb96184ee44de792eb5f5fc9713dc7.E_input', t85_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/4ebb96184ee44de792eb5f5fc9713dc7.E_input', t100_xboole_1)).
fof(d5_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)), file('data/MPTP/proofs/4ebb96184ee44de792eb5f5fc9713dc7.E_input', d5_subset_1)).
fof(involutiveness_k3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,k3_subset_1(X1,X2))=X2), file('data/MPTP/proofs/4ebb96184ee44de792eb5f5fc9713dc7.E_input', involutiveness_k3_subset_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/4ebb96184ee44de792eb5f5fc9713dc7.E_input', commutativity_k3_xboole_0)).
fof(c_0_17, plain, ![X5, X6, X7, X8, X9]:((~r1_tarski(X5,X6)|(~r2_hidden(X7,X5)|r2_hidden(X7,X6)))&((r2_hidden(esk1_2(X8,X9),X8)|r1_tarski(X8,X9))&(~r2_hidden(esk1_2(X8,X9),X9)|r1_tarski(X8,X9)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_18, plain, ![X11, X12]:(~r1_tarski(X11,X12)|k3_xboole_0(X11,X12)=X11), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_19, plain, (r1_tarski(X1,X2)|~r2_hidden(esk1_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20, plain, (r2_hidden(esk1_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_21, plain, ![X13, X14]:((~r1_xboole_0(X13,X14)|k3_xboole_0(X13,X14)=k1_xboole_0)&(k3_xboole_0(X13,X14)!=k1_xboole_0|r1_xboole_0(X13,X14))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_22, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_24, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_25, plain, (k3_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_22, c_0_23])).
fof(c_0_26, plain, ![X15, X16, X18, X19, X20]:(((r2_hidden(esk2_2(X15,X16),X15)|r1_xboole_0(X15,X16))&(r2_hidden(esk2_2(X15,X16),X16)|r1_xboole_0(X15,X16)))&(~r2_hidden(X20,X18)|~r2_hidden(X20,X19)|~r1_xboole_0(X18,X19))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_27, plain, (r1_xboole_0(X1,X1)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_28, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_29, plain, (r1_xboole_0(k1_xboole_0,k1_xboole_0)), inference(er,[status(thm)],[c_0_27])).
fof(c_0_30, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(r1_tarski(X2,k3_subset_1(X1,X2))<=>X2=k1_subset_1(X1)))), inference(assume_negation,[status(cth)],[t38_subset_1])).
cnf(c_0_31, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_32, plain, (r2_hidden(esk2_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_33, plain, ![X23, X24, X25, X26, X27, X28]:(((~r2_hidden(X25,X24)|r1_tarski(X25,X23)|X24!=k1_zfmisc_1(X23))&(~r1_tarski(X26,X23)|r2_hidden(X26,X24)|X24!=k1_zfmisc_1(X23)))&((~r2_hidden(esk5_2(X27,X28),X28)|~r1_tarski(esk5_2(X27,X28),X27)|X28=k1_zfmisc_1(X27))&(r2_hidden(esk5_2(X27,X28),X28)|r1_tarski(esk5_2(X27,X28),X27)|X28=k1_zfmisc_1(X27)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
fof(c_0_34, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(esk3_0))&((~r1_tarski(esk4_0,k3_subset_1(esk3_0,esk4_0))|esk4_0!=k1_subset_1(esk3_0))&(r1_tarski(esk4_0,k3_subset_1(esk3_0,esk4_0))|esk4_0=k1_subset_1(esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])).
fof(c_0_35, plain, ![X30]:k1_subset_1(X30)=k1_xboole_0, inference(variable_rename,[status(thm)],[d3_subset_1])).
fof(c_0_36, plain, ![X31, X32, X33, X34, X35, X36, X37, X38]:((((r2_hidden(X34,X31)|~r2_hidden(X34,X33)|X33!=k3_xboole_0(X31,X32))&(r2_hidden(X34,X32)|~r2_hidden(X34,X33)|X33!=k3_xboole_0(X31,X32)))&(~r2_hidden(X35,X31)|~r2_hidden(X35,X32)|r2_hidden(X35,X33)|X33!=k3_xboole_0(X31,X32)))&((~r2_hidden(esk6_3(X36,X37,X38),X38)|(~r2_hidden(esk6_3(X36,X37,X38),X36)|~r2_hidden(esk6_3(X36,X37,X38),X37))|X38=k3_xboole_0(X36,X37))&((r2_hidden(esk6_3(X36,X37,X38),X36)|r2_hidden(esk6_3(X36,X37,X38),X38)|X38=k3_xboole_0(X36,X37))&(r2_hidden(esk6_3(X36,X37,X38),X37)|r2_hidden(esk6_3(X36,X37,X38),X38)|X38=k3_xboole_0(X36,X37))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_37, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_38, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
fof(c_0_39, plain, ![X40, X41]:(~m1_subset_1(X41,k1_zfmisc_1(X40))|m1_subset_1(k3_subset_1(X40,X41),k1_zfmisc_1(X40))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_subset_1])])).
cnf(c_0_40, plain, (r2_hidden(X1,X3)|~r1_tarski(X1,X2)|X3!=k1_zfmisc_1(X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_41, negated_conjecture, (r1_tarski(esk4_0,k3_subset_1(esk3_0,esk4_0))|esk4_0=k1_subset_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42, plain, (k1_subset_1(X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_43, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_37, c_0_38])).
fof(c_0_45, plain, ![X42, X43]:(((~m1_subset_1(X43,X42)|r2_hidden(X43,X42)|v1_xboole_0(X42))&(~r2_hidden(X43,X42)|m1_subset_1(X43,X42)|v1_xboole_0(X42)))&((~m1_subset_1(X43,X42)|v1_xboole_0(X43)|~v1_xboole_0(X42))&(~v1_xboole_0(X43)|m1_subset_1(X43,X42)|~v1_xboole_0(X42)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
cnf(c_0_46, plain, (m1_subset_1(k3_subset_1(X2,X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_47, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_48, plain, ![X44]:~v1_xboole_0(k1_zfmisc_1(X44)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
cnf(c_0_49, plain, (r2_hidden(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_50, negated_conjecture, (esk4_0=k1_xboole_0|r1_tarski(esk4_0,k3_subset_1(esk3_0,esk4_0))), inference(rw,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_51, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_31])).
cnf(c_0_52, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_53, negated_conjecture, (m1_subset_1(k3_subset_1(esk3_0,esk4_0),k1_zfmisc_1(esk3_0))), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_54, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_55, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_56, negated_conjecture, (k1_xboole_0=esk4_0|r2_hidden(esk4_0,k1_zfmisc_1(k3_subset_1(esk3_0,esk4_0)))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_57, plain, (r1_tarski(X1,X2)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_51, c_0_20])).
cnf(c_0_58, negated_conjecture, (r2_hidden(k3_subset_1(esk3_0,esk4_0),k1_zfmisc_1(esk3_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54])).
cnf(c_0_59, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_60, plain, ![X45]:(X45=k1_xboole_0|r2_hidden(esk7_1(X45),X45)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_61, negated_conjecture, (~r1_tarski(esk4_0,k3_subset_1(esk3_0,esk4_0))|esk4_0!=k1_subset_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_62, negated_conjecture, (r1_tarski(esk4_0,X1)|k1_zfmisc_1(k3_subset_1(esk3_0,esk4_0))!=k1_zfmisc_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57])).
fof(c_0_63, plain, ![X47, X48, X49]:(~r1_tarski(X47,X48)|r1_xboole_0(X47,k4_xboole_0(X49,X48))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t85_xboole_1])])).
fof(c_0_64, plain, ![X50, X51]:k4_xboole_0(X50,X51)=k5_xboole_0(X50,k3_xboole_0(X50,X51)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_65, plain, ![X52, X53]:(~m1_subset_1(X53,k1_zfmisc_1(X52))|k3_subset_1(X52,X53)=k4_xboole_0(X52,X53)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
fof(c_0_66, plain, ![X54, X55]:(~m1_subset_1(X55,k1_zfmisc_1(X54))|k3_subset_1(X54,k3_subset_1(X54,X55))=X55), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k3_subset_1])])).
cnf(c_0_67, negated_conjecture, (r1_tarski(k3_subset_1(esk3_0,esk4_0),X1)|k1_zfmisc_1(esk3_0)!=k1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_55, c_0_58])).
fof(c_0_68, plain, ![X56, X57]:k3_xboole_0(X56,X57)=k3_xboole_0(X57,X56), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_69, negated_conjecture, (r2_hidden(X1,k3_subset_1(esk3_0,esk4_0))|~r2_hidden(X1,esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_50]), c_0_51])).
cnf(c_0_70, plain, (X1=k1_xboole_0|r2_hidden(esk7_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_71, negated_conjecture, (esk4_0!=k1_xboole_0|~r1_tarski(esk4_0,k3_subset_1(esk3_0,esk4_0))), inference(rw,[status(thm)],[c_0_61, c_0_42])).
cnf(c_0_72, negated_conjecture, (r1_tarski(esk4_0,k3_subset_1(esk3_0,esk4_0))), inference(er,[status(thm)],[c_0_62])).
cnf(c_0_73, plain, (r1_xboole_0(X1,k4_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_74, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_75, plain, (k3_subset_1(X2,X1)=k4_xboole_0(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_76, plain, (k3_subset_1(X2,k3_subset_1(X2,X1))=X1|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_77, negated_conjecture, (r1_tarski(k3_subset_1(esk3_0,esk4_0),esk3_0)), inference(er,[status(thm)],[c_0_67])).
cnf(c_0_78, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_79, negated_conjecture, (k1_xboole_0=esk4_0|r2_hidden(esk7_1(esk4_0),k3_subset_1(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_80, negated_conjecture, (k1_xboole_0!=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_72])])).
cnf(c_0_81, plain, (r1_xboole_0(X1,k5_xboole_0(X3,k3_xboole_0(X3,X2)))|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_82, plain, (k5_xboole_0(X2,k3_xboole_0(X2,X1))=k3_subset_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_75, c_0_74])).
cnf(c_0_83, negated_conjecture, (k3_subset_1(esk3_0,k3_subset_1(esk3_0,esk4_0))=esk4_0), inference(spm,[status(thm)],[c_0_76, c_0_47])).
cnf(c_0_84, negated_conjecture, (k3_xboole_0(esk3_0,k3_subset_1(esk3_0,esk4_0))=k3_subset_1(esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_77]), c_0_78])).
cnf(c_0_85, plain, (X1=k1_xboole_0|~r1_xboole_0(X2,X1)|~r2_hidden(esk7_1(X1),X2)), inference(spm,[status(thm)],[c_0_28, c_0_70])).
cnf(c_0_86, negated_conjecture, (r2_hidden(esk7_1(esk4_0),k3_subset_1(esk3_0,esk4_0))), inference(sr,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_87, plain, (r1_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1)))), inference(spm,[status(thm)],[c_0_81, c_0_23])).
cnf(c_0_88, negated_conjecture, (k5_xboole_0(esk3_0,k3_subset_1(esk3_0,esk4_0))=esk4_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_53]), c_0_83]), c_0_84])).
cnf(c_0_89, negated_conjecture, (~r1_xboole_0(k3_subset_1(esk3_0,esk4_0),esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_80])).
cnf(c_0_90, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_84]), c_0_88]), c_0_89]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 91
# Proof object clause steps            : 56
# Proof object formula steps           : 35
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 26
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 33
# Processed clauses                    : 1238
# ...of these trivial                  : 34
# ...subsumed                          : 608
# ...remaining for further processing  : 596
# Other redundant clauses eliminated   : 89
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 29
# Backward-rewritten                   : 5
# Generated clauses                    : 3542
# ...of the previous two non-trivial   : 3089
# Contextual simplify-reflections      : 2
# Paramodulations                      : 3342
# Factorizations                       : 4
# Equation resolutions                 : 183
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 549
#    Positive orientable unit clauses  : 81
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 37
#    Non-unit-clauses                  : 430
# Current number of unprocessed clauses: 1850
# ...number of literals in the above   : 6569
# Current number of archived formulas  : 0
# Current number of archived clauses   : 49
# Clause-clause subsumption calls (NU) : 21253
# Rec. Clause-clause subsumption calls : 15917
# Non-unit clause-clause subsumptions  : 355
# Unit Clause-clause subsumption calls : 1023
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 68
# BW rewrite match successes           : 7
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 44304

# -------------------------------------------------
# User time                : 0.048 s
# System time              : 0.000 s
# Total time               : 0.048 s
# Maximum resident set size: 28644 pages
