# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t151_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(X1=k1_tsep_1(X1,X3,X4)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))=>((r2_funct_2(u1_struct_0(X3),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,k10_tmap_1(X1,X2,X3,X4,X5,X6)),X5)&(r2_funct_2(u1_struct_0(X4),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,k10_tmap_1(X1,X2,X3,X4,X5,X6)),X6)&r4_tsep_1(X1,X3,X4)))=>(v1_funct_1(k10_tmap_1(X1,X2,X3,X4,X5,X6))&(v1_funct_2(k10_tmap_1(X1,X2,X3,X4,X5,X6),u1_struct_0(X1),u1_struct_0(X2))&(v5_pre_topc(k10_tmap_1(X1,X2,X3,X4,X5,X6),X1,X2)&m1_subset_1(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))))))))))), file('data/MPTP/proofs/4ec1dcf8e31c475b857e6bc181a9e57a.E_input', t151_tmap_1)).
fof(redefinition_r2_funct_2, axiom, ![X1, X2, X3, X4]:(((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&(v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))=>(r2_funct_2(X1,X2,X3,X4)<=>X3=X4)), file('data/MPTP/proofs/4ec1dcf8e31c475b857e6bc181a9e57a.E_input', redefinition_r2_funct_2)).
fof(dt_k1_tsep_1, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))&(~(v2_struct_0(X3))&m1_pre_topc(X3,X1))))=>(~(v2_struct_0(k1_tsep_1(X1,X2,X3)))&(v1_pre_topc(k1_tsep_1(X1,X2,X3))&m1_pre_topc(k1_tsep_1(X1,X2,X3),X1)))), file('data/MPTP/proofs/4ec1dcf8e31c475b857e6bc181a9e57a.E_input', dt_k1_tsep_1)).
fof(dt_k3_tmap_1, axiom, ![X1, X2, X3, X4, X5]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))&(m1_pre_topc(X3,X1)&(m1_pre_topc(X4,X1)&(v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))))))=>(v1_funct_1(k3_tmap_1(X1,X2,X3,X4,X5))&(v1_funct_2(k3_tmap_1(X1,X2,X3,X4,X5),u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(k3_tmap_1(X1,X2,X3,X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))), file('data/MPTP/proofs/4ec1dcf8e31c475b857e6bc181a9e57a.E_input', dt_k3_tmap_1)).
fof(dt_k10_tmap_1, axiom, ![X1, X2, X3, X4, X5, X6]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))&((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))&((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))&((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))&(v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))))))=>(v1_funct_1(k10_tmap_1(X1,X2,X3,X4,X5,X6))&(v1_funct_2(k10_tmap_1(X1,X2,X3,X4,X5,X6),u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&m1_subset_1(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))))), file('data/MPTP/proofs/4ec1dcf8e31c475b857e6bc181a9e57a.E_input', dt_k10_tmap_1)).
fof(t126_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(r4_tsep_1(X1,X3,X4)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))))=>((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&(v5_pre_topc(X5,k1_tsep_1(X1,X3,X4),X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2)))))))<=>((v1_funct_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5))&(v1_funct_2(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),X3,X2)&m1_subset_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))&(v1_funct_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5))&(v1_funct_2(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5),u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5),X4,X2)&m1_subset_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))))))))))), file('data/MPTP/proofs/4ec1dcf8e31c475b857e6bc181a9e57a.E_input', t126_tmap_1)).
fof(c_0_6, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(X1=k1_tsep_1(X1,X3,X4)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))=>((r2_funct_2(u1_struct_0(X3),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,k10_tmap_1(X1,X2,X3,X4,X5,X6)),X5)&(r2_funct_2(u1_struct_0(X4),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,k10_tmap_1(X1,X2,X3,X4,X5,X6)),X6)&r4_tsep_1(X1,X3,X4)))=>(v1_funct_1(k10_tmap_1(X1,X2,X3,X4,X5,X6))&(v1_funct_2(k10_tmap_1(X1,X2,X3,X4,X5,X6),u1_struct_0(X1),u1_struct_0(X2))&(v5_pre_topc(k10_tmap_1(X1,X2,X3,X4,X5,X6),X1,X2)&m1_subset_1(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))))))))))))), inference(assume_negation,[status(cth)],[t151_tmap_1])).
fof(c_0_7, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk1_0))&((~v2_struct_0(esk4_0)&m1_pre_topc(esk4_0,esk1_0))&(esk1_0=k1_tsep_1(esk1_0,esk3_0,esk4_0)&((v1_funct_1(esk5_0)&(v1_funct_2(esk5_0,u1_struct_0(esk3_0),u1_struct_0(esk2_0))&(v5_pre_topc(esk5_0,esk3_0,esk2_0)&m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0)))))))&((v1_funct_1(esk6_0)&(v1_funct_2(esk6_0,u1_struct_0(esk4_0),u1_struct_0(esk2_0))&(v5_pre_topc(esk6_0,esk4_0,esk2_0)&m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0)))))))&((r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk5_0)&(r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk6_0)&r4_tsep_1(esk1_0,esk3_0,esk4_0)))&(~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|(~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(esk1_0),u1_struct_0(esk2_0))|(~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk1_0,esk2_0)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))))))))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_6])])])])).
fof(c_0_8, plain, ![X13, X14, X15, X16]:((~r2_funct_2(X13,X14,X15,X16)|X15=X16|(~v1_funct_1(X15)|(~v1_funct_2(X15,X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))|(~v1_funct_1(X16)|(~v1_funct_2(X16,X13,X14)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X13,X14)))))))&(X15!=X16|r2_funct_2(X13,X14,X15,X16)|(~v1_funct_1(X15)|(~v1_funct_2(X15,X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))|(~v1_funct_1(X16)|(~v1_funct_2(X16,X13,X14)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X13,X14)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r2_funct_2])])])).
cnf(c_0_9, negated_conjecture, (r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk6_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10, negated_conjecture, (esk1_0=k1_tsep_1(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_11, plain, ![X17, X18, X19]:((~v2_struct_0(k1_tsep_1(X17,X18,X19))|(v2_struct_0(X17)|~l1_pre_topc(X17)|(v2_struct_0(X18)|~m1_pre_topc(X18,X17)|(v2_struct_0(X19)|~m1_pre_topc(X19,X17)))))&((v1_pre_topc(k1_tsep_1(X17,X18,X19))|(v2_struct_0(X17)|~l1_pre_topc(X17)|(v2_struct_0(X18)|~m1_pre_topc(X18,X17)|(v2_struct_0(X19)|~m1_pre_topc(X19,X17)))))&(m1_pre_topc(k1_tsep_1(X17,X18,X19),X17)|(v2_struct_0(X17)|~l1_pre_topc(X17)|(v2_struct_0(X18)|~m1_pre_topc(X18,X17)|(v2_struct_0(X19)|~m1_pre_topc(X19,X17))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_tsep_1])])])])).
fof(c_0_12, plain, ![X5, X4, X3, X2, X1]:(epred1_5(X1,X2,X3,X4,X5)<=>((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&(v5_pre_topc(X5,k1_tsep_1(X1,X3,X4),X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2)))))))<=>((v1_funct_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5))&(v1_funct_2(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),X3,X2)&m1_subset_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))&(v1_funct_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5))&(v1_funct_2(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5),u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5),X4,X2)&m1_subset_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))))), introduced(definition)).
cnf(c_0_13, plain, (X3=X4|~r2_funct_2(X1,X2,X3,X4)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X4)|~v1_funct_2(X4,X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, negated_conjecture, (r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,esk1_0,esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk6_0)), inference(rw,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_15, negated_conjecture, (m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_16, negated_conjecture, (v1_funct_2(esk6_0,u1_struct_0(esk4_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_17, negated_conjecture, (v1_funct_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_18, plain, ![X20, X21, X22, X23, X24]:((v1_funct_1(k3_tmap_1(X20,X21,X22,X23,X24))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20))|(v2_struct_0(X21)|(~v2_pre_topc(X21)|~l1_pre_topc(X21))|(~m1_pre_topc(X22,X20)|(~m1_pre_topc(X23,X20)|(~v1_funct_1(X24)|(~v1_funct_2(X24,u1_struct_0(X22),u1_struct_0(X21))|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(X21)))))))))))&((v1_funct_2(k3_tmap_1(X20,X21,X22,X23,X24),u1_struct_0(X23),u1_struct_0(X21))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20))|(v2_struct_0(X21)|(~v2_pre_topc(X21)|~l1_pre_topc(X21))|(~m1_pre_topc(X22,X20)|(~m1_pre_topc(X23,X20)|(~v1_funct_1(X24)|(~v1_funct_2(X24,u1_struct_0(X22),u1_struct_0(X21))|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(X21)))))))))))&(m1_subset_1(k3_tmap_1(X20,X21,X22,X23,X24),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X23),u1_struct_0(X21))))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20))|(v2_struct_0(X21)|(~v2_pre_topc(X21)|~l1_pre_topc(X21))|(~m1_pre_topc(X22,X20)|(~m1_pre_topc(X23,X20)|(~v1_funct_1(X24)|(~v1_funct_2(X24,u1_struct_0(X22),u1_struct_0(X21))|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(X21))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_tmap_1])])])])).
cnf(c_0_19, plain, (m1_pre_topc(k1_tsep_1(X1,X2,X3),X1)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20, negated_conjecture, (m1_pre_topc(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_21, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_22, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_23, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_24, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_25, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_26, plain, ![X25, X26, X27, X28, X29, X30]:((v1_funct_1(k10_tmap_1(X25,X26,X27,X28,X29,X30))|(v2_struct_0(X25)|(~v2_pre_topc(X25)|~l1_pre_topc(X25))|(v2_struct_0(X26)|(~v2_pre_topc(X26)|~l1_pre_topc(X26))|(v2_struct_0(X27)|~m1_pre_topc(X27,X25)|(v2_struct_0(X28)|~m1_pre_topc(X28,X25)|(~v1_funct_1(X29)|(~v1_funct_2(X29,u1_struct_0(X27),u1_struct_0(X26))|~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X27),u1_struct_0(X26)))))|(~v1_funct_1(X30)|(~v1_funct_2(X30,u1_struct_0(X28),u1_struct_0(X26))|~m1_subset_1(X30,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X28),u1_struct_0(X26))))))))))))&((v1_funct_2(k10_tmap_1(X25,X26,X27,X28,X29,X30),u1_struct_0(k1_tsep_1(X25,X27,X28)),u1_struct_0(X26))|(v2_struct_0(X25)|(~v2_pre_topc(X25)|~l1_pre_topc(X25))|(v2_struct_0(X26)|(~v2_pre_topc(X26)|~l1_pre_topc(X26))|(v2_struct_0(X27)|~m1_pre_topc(X27,X25)|(v2_struct_0(X28)|~m1_pre_topc(X28,X25)|(~v1_funct_1(X29)|(~v1_funct_2(X29,u1_struct_0(X27),u1_struct_0(X26))|~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X27),u1_struct_0(X26)))))|(~v1_funct_1(X30)|(~v1_funct_2(X30,u1_struct_0(X28),u1_struct_0(X26))|~m1_subset_1(X30,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X28),u1_struct_0(X26))))))))))))&(m1_subset_1(k10_tmap_1(X25,X26,X27,X28,X29,X30),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X25,X27,X28)),u1_struct_0(X26))))|(v2_struct_0(X25)|(~v2_pre_topc(X25)|~l1_pre_topc(X25))|(v2_struct_0(X26)|(~v2_pre_topc(X26)|~l1_pre_topc(X26))|(v2_struct_0(X27)|~m1_pre_topc(X27,X25)|(v2_struct_0(X28)|~m1_pre_topc(X28,X25)|(~v1_funct_1(X29)|(~v1_funct_2(X29,u1_struct_0(X27),u1_struct_0(X26))|~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X27),u1_struct_0(X26)))))|(~v1_funct_1(X30)|(~v1_funct_2(X30,u1_struct_0(X28),u1_struct_0(X26))|~m1_subset_1(X30,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X28),u1_struct_0(X26)))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k10_tmap_1])])])])).
cnf(c_0_27, negated_conjecture, (r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk5_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_28, plain, ![X5, X4, X3, X2, X1]:(epred1_5(X1,X2,X3,X4,X5)=>((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&(v5_pre_topc(X5,k1_tsep_1(X1,X3,X4),X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2)))))))<=>((v1_funct_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5))&(v1_funct_2(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),X3,X2)&m1_subset_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))&(v1_funct_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5))&(v1_funct_2(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5),u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5),X4,X2)&m1_subset_1(k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))))), inference(split_equiv,[status(thm)],[c_0_12])).
fof(c_0_29, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(r4_tsep_1(X1,X3,X4)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))))=>epred1_5(X1,X2,X3,X4,X5))))))), inference(apply_def,[status(thm)],[t126_tmap_1, c_0_12])).
cnf(c_0_30, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,esk1_0,esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk6_0|~m1_subset_1(k3_tmap_1(esk1_0,esk2_0,esk1_0,esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~v1_funct_2(k3_tmap_1(esk1_0,esk2_0,esk1_0,esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,esk1_0,esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16]), c_0_17])])).
cnf(c_0_31, plain, (m1_subset_1(k3_tmap_1(X1,X2,X3,X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_32, negated_conjecture, (m1_pre_topc(esk1_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_10]), c_0_20]), c_0_21]), c_0_22])]), c_0_23]), c_0_24]), c_0_25])).
cnf(c_0_33, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_34, negated_conjecture, (v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_35, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_36, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_37, plain, (m1_subset_1(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_1(X6)|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_38, plain, (v1_funct_2(k10_tmap_1(X1,X2,X3,X4,X5,X6),u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_1(X6)|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_39, negated_conjecture, (r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk5_0)), inference(rw,[status(thm)],[c_0_27, c_0_10])).
cnf(c_0_40, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_41, negated_conjecture, (v1_funct_2(esk5_0,u1_struct_0(esk3_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_42, negated_conjecture, (v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_43, plain, ![X36, X37, X38, X39, X40]:((((v1_funct_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36))|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(k1_tsep_1(X40,X38,X37)),u1_struct_0(X39))|(~v5_pre_topc(X36,k1_tsep_1(X40,X38,X37),X39)|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X40,X38,X37)),u1_struct_0(X39)))))))|~epred1_5(X40,X39,X38,X37,X36))&((v1_funct_2(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36),u1_struct_0(X38),u1_struct_0(X39))|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(k1_tsep_1(X40,X38,X37)),u1_struct_0(X39))|(~v5_pre_topc(X36,k1_tsep_1(X40,X38,X37),X39)|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X40,X38,X37)),u1_struct_0(X39)))))))|~epred1_5(X40,X39,X38,X37,X36))&((v5_pre_topc(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36),X38,X39)|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(k1_tsep_1(X40,X38,X37)),u1_struct_0(X39))|(~v5_pre_topc(X36,k1_tsep_1(X40,X38,X37),X39)|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X40,X38,X37)),u1_struct_0(X39)))))))|~epred1_5(X40,X39,X38,X37,X36))&(m1_subset_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X38),u1_struct_0(X39))))|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(k1_tsep_1(X40,X38,X37)),u1_struct_0(X39))|(~v5_pre_topc(X36,k1_tsep_1(X40,X38,X37),X39)|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X40,X38,X37)),u1_struct_0(X39)))))))|~epred1_5(X40,X39,X38,X37,X36)))))&((v1_funct_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36))|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(k1_tsep_1(X40,X38,X37)),u1_struct_0(X39))|(~v5_pre_topc(X36,k1_tsep_1(X40,X38,X37),X39)|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X40,X38,X37)),u1_struct_0(X39)))))))|~epred1_5(X40,X39,X38,X37,X36))&((v1_funct_2(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36),u1_struct_0(X37),u1_struct_0(X39))|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(k1_tsep_1(X40,X38,X37)),u1_struct_0(X39))|(~v5_pre_topc(X36,k1_tsep_1(X40,X38,X37),X39)|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X40,X38,X37)),u1_struct_0(X39)))))))|~epred1_5(X40,X39,X38,X37,X36))&((v5_pre_topc(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36),X37,X39)|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(k1_tsep_1(X40,X38,X37)),u1_struct_0(X39))|(~v5_pre_topc(X36,k1_tsep_1(X40,X38,X37),X39)|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X40,X38,X37)),u1_struct_0(X39)))))))|~epred1_5(X40,X39,X38,X37,X36))&(m1_subset_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X37),u1_struct_0(X39))))|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(k1_tsep_1(X40,X38,X37)),u1_struct_0(X39))|(~v5_pre_topc(X36,k1_tsep_1(X40,X38,X37),X39)|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X40,X38,X37)),u1_struct_0(X39)))))))|~epred1_5(X40,X39,X38,X37,X36))))))&((v1_funct_1(X36)|(~v1_funct_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36))|(~v1_funct_2(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36),u1_struct_0(X38),u1_struct_0(X39))|(~v5_pre_topc(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36),X38,X39)|~m1_subset_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X38),u1_struct_0(X39))))))|(~v1_funct_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36))|(~v1_funct_2(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36),u1_struct_0(X37),u1_struct_0(X39))|(~v5_pre_topc(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36),X37,X39)|~m1_subset_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X37),u1_struct_0(X39))))))))|~epred1_5(X40,X39,X38,X37,X36))&((v1_funct_2(X36,u1_struct_0(k1_tsep_1(X40,X38,X37)),u1_struct_0(X39))|(~v1_funct_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36))|(~v1_funct_2(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36),u1_struct_0(X38),u1_struct_0(X39))|(~v5_pre_topc(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36),X38,X39)|~m1_subset_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X38),u1_struct_0(X39))))))|(~v1_funct_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36))|(~v1_funct_2(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36),u1_struct_0(X37),u1_struct_0(X39))|(~v5_pre_topc(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36),X37,X39)|~m1_subset_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X37),u1_struct_0(X39))))))))|~epred1_5(X40,X39,X38,X37,X36))&((v5_pre_topc(X36,k1_tsep_1(X40,X38,X37),X39)|(~v1_funct_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36))|(~v1_funct_2(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36),u1_struct_0(X38),u1_struct_0(X39))|(~v5_pre_topc(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36),X38,X39)|~m1_subset_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X38),u1_struct_0(X39))))))|(~v1_funct_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36))|(~v1_funct_2(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36),u1_struct_0(X37),u1_struct_0(X39))|(~v5_pre_topc(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36),X37,X39)|~m1_subset_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X37),u1_struct_0(X39))))))))|~epred1_5(X40,X39,X38,X37,X36))&(m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X40,X38,X37)),u1_struct_0(X39))))|(~v1_funct_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36))|(~v1_funct_2(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36),u1_struct_0(X38),u1_struct_0(X39))|(~v5_pre_topc(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36),X38,X39)|~m1_subset_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X38,X36),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X38),u1_struct_0(X39))))))|(~v1_funct_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36))|(~v1_funct_2(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36),u1_struct_0(X37),u1_struct_0(X39))|(~v5_pre_topc(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36),X37,X39)|~m1_subset_1(k3_tmap_1(X40,X39,k1_tsep_1(X40,X38,X37),X37,X36),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X37),u1_struct_0(X39))))))))|~epred1_5(X40,X39,X38,X37,X36)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_28])])])).
fof(c_0_44, plain, ![X31, X32, X33, X34, X35]:(v2_struct_0(X31)|(~v2_pre_topc(X31)|~l1_pre_topc(X31))|(v2_struct_0(X32)|(~v2_pre_topc(X32)|~l1_pre_topc(X32))|(v2_struct_0(X33)|~m1_pre_topc(X33,X31)|(v2_struct_0(X34)|~m1_pre_topc(X34,X31)|(~r4_tsep_1(X31,X33,X34)|(~v1_funct_1(X35)|(~v1_funct_2(X35,u1_struct_0(k1_tsep_1(X31,X33,X34)),u1_struct_0(X32))|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X31,X33,X34)),u1_struct_0(X32)))))|epred1_5(X31,X32,X33,X34,X35))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_29])])])])).
cnf(c_0_45, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,esk1_0,esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk6_0|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))|~v1_funct_2(k3_tmap_1(esk1_0,esk2_0,esk1_0,esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,esk1_0,esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_20]), c_0_32]), c_0_33]), c_0_22]), c_0_34]), c_0_35])]), c_0_36]), c_0_25])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(k10_tmap_1(esk1_0,X1,esk3_0,esk4_0,X2,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(esk4_0),u1_struct_0(X1))|~v1_funct_2(X2,u1_struct_0(esk3_0),u1_struct_0(X1))|~v1_funct_1(X3)|~v1_funct_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_10]), c_0_20]), c_0_21]), c_0_22]), c_0_35])]), c_0_23]), c_0_24]), c_0_25])).
cnf(c_0_47, plain, (v1_funct_2(k3_tmap_1(X1,X2,X3,X4,X5),u1_struct_0(X4),u1_struct_0(X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_48, negated_conjecture, (v1_funct_2(k10_tmap_1(esk1_0,X1,esk3_0,esk4_0,X2,X3),u1_struct_0(esk1_0),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(esk4_0),u1_struct_0(X1))|~v1_funct_2(X2,u1_struct_0(esk3_0),u1_struct_0(X1))|~v1_funct_1(X3)|~v1_funct_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_10]), c_0_20]), c_0_21]), c_0_22]), c_0_35])]), c_0_23]), c_0_24]), c_0_25])).
cnf(c_0_49, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk5_0|~m1_subset_1(k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))|~v1_funct_2(k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_39]), c_0_40]), c_0_41]), c_0_42])])).
cnf(c_0_50, plain, (v5_pre_topc(X1,k1_tsep_1(X2,X3,X4),X5)|~v1_funct_1(k3_tmap_1(X2,X5,k1_tsep_1(X2,X3,X4),X3,X1))|~v1_funct_2(k3_tmap_1(X2,X5,k1_tsep_1(X2,X3,X4),X3,X1),u1_struct_0(X3),u1_struct_0(X5))|~v5_pre_topc(k3_tmap_1(X2,X5,k1_tsep_1(X2,X3,X4),X3,X1),X3,X5)|~m1_subset_1(k3_tmap_1(X2,X5,k1_tsep_1(X2,X3,X4),X3,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X5))))|~v1_funct_1(k3_tmap_1(X2,X5,k1_tsep_1(X2,X3,X4),X4,X1))|~v1_funct_2(k3_tmap_1(X2,X5,k1_tsep_1(X2,X3,X4),X4,X1),u1_struct_0(X4),u1_struct_0(X5))|~v5_pre_topc(k3_tmap_1(X2,X5,k1_tsep_1(X2,X3,X4),X4,X1),X4,X5)|~m1_subset_1(k3_tmap_1(X2,X5,k1_tsep_1(X2,X3,X4),X4,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X5))))|~epred1_5(X2,X5,X3,X4,X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_51, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|epred1_5(X1,X2,X3,X4,X5)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~r4_tsep_1(X1,X3,X4)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_52, negated_conjecture, (r4_tsep_1(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_53, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,esk1_0,esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk6_0|~v1_funct_2(k3_tmap_1(esk1_0,esk2_0,esk1_0,esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,esk1_0,esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_15]), c_0_40]), c_0_16]), c_0_41]), c_0_17]), c_0_42]), c_0_33]), c_0_34])]), c_0_36])).
cnf(c_0_54, negated_conjecture, (v1_funct_2(k3_tmap_1(X1,X2,esk1_0,X3,k10_tmap_1(esk1_0,X2,esk3_0,esk4_0,X4,X5)),u1_struct_0(X3),u1_struct_0(X2))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(esk4_0),u1_struct_0(X2))|~v1_funct_2(X4,u1_struct_0(esk3_0),u1_struct_0(X2))|~v1_funct_1(k10_tmap_1(esk1_0,X2,esk3_0,esk4_0,X4,X5))|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_pre_topc(esk1_0,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_46]), c_0_48])).
cnf(c_0_55, plain, (v1_funct_1(k3_tmap_1(X1,X2,X3,X4,X5))|v2_struct_0(X1)|v2_struct_0(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_56, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk5_0|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))|~v1_funct_2(k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_31]), c_0_21]), c_0_32]), c_0_33]), c_0_22]), c_0_34]), c_0_35])]), c_0_36]), c_0_25])).
cnf(c_0_57, negated_conjecture, (v5_pre_topc(X1,esk1_0,X2)|~epred1_5(esk1_0,X2,esk3_0,esk4_0,X1)|~m1_subset_1(k3_tmap_1(esk1_0,X2,esk1_0,esk4_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(X2))))|~m1_subset_1(k3_tmap_1(esk1_0,X2,esk1_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(X2))))|~v5_pre_topc(k3_tmap_1(esk1_0,X2,esk1_0,esk4_0,X1),esk4_0,X2)|~v5_pre_topc(k3_tmap_1(esk1_0,X2,esk1_0,esk3_0,X1),esk3_0,X2)|~v1_funct_2(k3_tmap_1(esk1_0,X2,esk1_0,esk4_0,X1),u1_struct_0(esk4_0),u1_struct_0(X2))|~v1_funct_2(k3_tmap_1(esk1_0,X2,esk1_0,esk3_0,X1),u1_struct_0(esk3_0),u1_struct_0(X2))|~v1_funct_1(k3_tmap_1(esk1_0,X2,esk1_0,esk4_0,X1))|~v1_funct_1(k3_tmap_1(esk1_0,X2,esk1_0,esk3_0,X1))), inference(spm,[status(thm)],[c_0_50, c_0_10])).
cnf(c_0_58, negated_conjecture, (epred1_5(esk1_0,X1,esk3_0,esk4_0,X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(esk1_0),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_10]), c_0_52]), c_0_20]), c_0_21]), c_0_22]), c_0_35])]), c_0_23]), c_0_24]), c_0_25])).
cnf(c_0_59, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,esk1_0,esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk6_0|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,esk1_0,esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_15]), c_0_40]), c_0_16]), c_0_41]), c_0_17]), c_0_42]), c_0_32]), c_0_20]), c_0_33]), c_0_22]), c_0_34]), c_0_35])]), c_0_25]), c_0_36])).
cnf(c_0_60, negated_conjecture, (v1_funct_1(k3_tmap_1(X1,X2,esk1_0,X3,k10_tmap_1(esk1_0,X2,esk3_0,esk4_0,X4,X5)))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(esk4_0),u1_struct_0(X2))|~v1_funct_2(X4,u1_struct_0(esk3_0),u1_struct_0(X2))|~v1_funct_1(k10_tmap_1(esk1_0,X2,esk3_0,esk4_0,X4,X5))|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_pre_topc(esk1_0,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_46]), c_0_48])).
cnf(c_0_61, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk5_0|~v1_funct_2(k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_46]), c_0_15]), c_0_40]), c_0_16]), c_0_41]), c_0_17]), c_0_42]), c_0_33]), c_0_34])]), c_0_36])).
cnf(c_0_62, negated_conjecture, (v5_pre_topc(X1,esk1_0,X2)|v2_struct_0(X2)|~m1_subset_1(k3_tmap_1(esk1_0,X2,esk1_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(X2))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X2))))|~v5_pre_topc(k3_tmap_1(esk1_0,X2,esk1_0,esk4_0,X1),esk4_0,X2)|~v5_pre_topc(k3_tmap_1(esk1_0,X2,esk1_0,esk3_0,X1),esk3_0,X2)|~v1_funct_2(k3_tmap_1(esk1_0,X2,esk1_0,esk4_0,X1),u1_struct_0(esk4_0),u1_struct_0(X2))|~v1_funct_2(k3_tmap_1(esk1_0,X2,esk1_0,esk3_0,X1),u1_struct_0(esk3_0),u1_struct_0(X2))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(X2))|~v1_funct_1(k3_tmap_1(esk1_0,X2,esk1_0,esk4_0,X1))|~v1_funct_1(k3_tmap_1(esk1_0,X2,esk1_0,esk3_0,X1))|~v1_funct_1(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_31]), c_0_20]), c_0_32]), c_0_22]), c_0_35])]), c_0_25]), c_0_58])).
cnf(c_0_63, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,esk1_0,esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk6_0|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_15]), c_0_40]), c_0_16]), c_0_41]), c_0_17]), c_0_42]), c_0_32]), c_0_20]), c_0_33]), c_0_22]), c_0_34]), c_0_35])]), c_0_25]), c_0_36])).
cnf(c_0_64, negated_conjecture, (~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk1_0,esk2_0)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_65, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk5_0|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_54]), c_0_15]), c_0_40]), c_0_16]), c_0_41]), c_0_17]), c_0_42]), c_0_32]), c_0_21]), c_0_33]), c_0_22]), c_0_34]), c_0_35])]), c_0_25]), c_0_36])).
cnf(c_0_66, negated_conjecture, (v5_pre_topc(X1,esk1_0,X2)|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X2))))|~v5_pre_topc(k3_tmap_1(esk1_0,X2,esk1_0,esk4_0,X1),esk4_0,X2)|~v5_pre_topc(k3_tmap_1(esk1_0,X2,esk1_0,esk3_0,X1),esk3_0,X2)|~v1_funct_2(k3_tmap_1(esk1_0,X2,esk1_0,esk4_0,X1),u1_struct_0(esk4_0),u1_struct_0(X2))|~v1_funct_2(k3_tmap_1(esk1_0,X2,esk1_0,esk3_0,X1),u1_struct_0(esk3_0),u1_struct_0(X2))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(X2))|~v1_funct_1(k3_tmap_1(esk1_0,X2,esk1_0,esk4_0,X1))|~v1_funct_1(k3_tmap_1(esk1_0,X2,esk1_0,esk3_0,X1))|~v1_funct_1(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_31]), c_0_21]), c_0_32]), c_0_22]), c_0_35])]), c_0_25])).
cnf(c_0_67, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,esk1_0,esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk6_0|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_48]), c_0_15]), c_0_40]), c_0_16]), c_0_41]), c_0_17]), c_0_42]), c_0_33]), c_0_34])]), c_0_36])).
cnf(c_0_68, negated_conjecture, (v5_pre_topc(esk6_0,esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_69, negated_conjecture, (~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk1_0,esk2_0)|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_46]), c_0_15]), c_0_40]), c_0_16]), c_0_41]), c_0_17]), c_0_42]), c_0_33]), c_0_34])]), c_0_36])).
cnf(c_0_70, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk5_0|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_60]), c_0_15]), c_0_40]), c_0_16]), c_0_41]), c_0_17]), c_0_42]), c_0_32]), c_0_21]), c_0_33]), c_0_22]), c_0_34]), c_0_35])]), c_0_25]), c_0_36])).
cnf(c_0_71, negated_conjecture, (~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))|~v5_pre_topc(k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk3_0,esk2_0)|~v1_funct_2(k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68]), c_0_16]), c_0_17]), c_0_33]), c_0_34])]), c_0_36]), c_0_69])).
cnf(c_0_72, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))=esk5_0|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_48]), c_0_15]), c_0_40]), c_0_16]), c_0_41]), c_0_17]), c_0_42]), c_0_33]), c_0_34])]), c_0_36])).
cnf(c_0_73, negated_conjecture, (v5_pre_topc(esk5_0,esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_74, negated_conjecture, (~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73]), c_0_41]), c_0_42])])).
cnf(c_0_75, negated_conjecture, (~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_46]), c_0_15]), c_0_40]), c_0_16]), c_0_41]), c_0_17]), c_0_42]), c_0_33]), c_0_34])]), c_0_36])).
cnf(c_0_76, negated_conjecture, (~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_48]), c_0_15]), c_0_40]), c_0_16]), c_0_41]), c_0_17]), c_0_42]), c_0_33]), c_0_34])]), c_0_36])).
cnf(c_0_77, plain, (v1_funct_1(k10_tmap_1(X1,X2,X3,X4,X5,X6))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_1(X6)|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_78, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_15]), c_0_40]), c_0_16]), c_0_41]), c_0_17]), c_0_42]), c_0_20]), c_0_21]), c_0_33]), c_0_22]), c_0_34]), c_0_35])]), c_0_23]), c_0_24]), c_0_36]), c_0_25]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 79
# Proof object clause steps            : 62
# Proof object formula steps           : 17
# Proof object conjectures             : 55
# Proof object clause conjectures      : 52
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 33
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 234
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 47
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 47
# Processed clauses                    : 170
# ...of these trivial                  : 0
# ...subsumed                          : 6
# ...remaining for further processing  : 164
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 13
# Backward-rewritten                   : 2
# Generated clauses                    : 149
# ...of the previous two non-trivial   : 117
# Contextual simplify-reflections      : 39
# Paramodulations                      : 148
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 101
#    Positive orientable unit clauses  : 22
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 74
# Current number of unprocessed clauses: 39
# ...number of literals in the above   : 582
# Current number of archived formulas  : 0
# Current number of archived clauses   : 62
# Clause-clause subsumption calls (NU) : 5229
# Rec. Clause-clause subsumption calls : 215
# Non-unit clause-clause subsumptions  : 58
# Unit Clause-clause subsumption calls : 16
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 18238

# -------------------------------------------------
# User time                : 0.019 s
# System time              : 0.003 s
# Total time               : 0.022 s
# Maximum resident set size: 28708 pages
