# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t16_funct_2, conjecture, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(![X4]:~((r2_hidden(X4,X2)&![X5]:~((r2_hidden(X5,X1)&X4=k1_funct_1(X3,X5)))))=>k2_relset_1(X1,X2,X3)=X2)), file('data/MPTP/proofs/4f43e9cec9bd4697bba53b13eb046233.E_input', t16_funct_2)).
fof(d5_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4))))), file('data/MPTP/proofs/4f43e9cec9bd4697bba53b13eb046233.E_input', d5_funct_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/4f43e9cec9bd4697bba53b13eb046233.E_input', cc1_relset_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/4f43e9cec9bd4697bba53b13eb046233.E_input', d3_tarski)).
fof(d19_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(v5_relat_1(X2,X1)<=>r1_tarski(k2_relat_1(X2),X1))), file('data/MPTP/proofs/4f43e9cec9bd4697bba53b13eb046233.E_input', d19_relat_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/4f43e9cec9bd4697bba53b13eb046233.E_input', d1_xboole_0)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(((X2=k1_xboole_0=>X1=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X2,X3)))&(X2=k1_xboole_0=>(X1=k1_xboole_0|(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0))))), file('data/MPTP/proofs/4f43e9cec9bd4697bba53b13eb046233.E_input', d1_funct_2)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/4f43e9cec9bd4697bba53b13eb046233.E_input', redefinition_k1_relset_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('data/MPTP/proofs/4f43e9cec9bd4697bba53b13eb046233.E_input', cc2_relset_1)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/4f43e9cec9bd4697bba53b13eb046233.E_input', redefinition_k2_relset_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/4f43e9cec9bd4697bba53b13eb046233.E_input', fc1_xboole_0)).
fof(c_0_11, negated_conjecture, ~(![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(![X4]:~((r2_hidden(X4,X2)&![X5]:~((r2_hidden(X5,X1)&X4=k1_funct_1(X3,X5)))))=>k2_relset_1(X1,X2,X3)=X2))), inference(assume_negation,[status(cth)],[t16_funct_2])).
fof(c_0_12, plain, ![X14, X15, X16, X18, X19, X20, X22]:((((r2_hidden(esk5_3(X14,X15,X16),k1_relat_1(X14))|~r2_hidden(X16,X15)|X15!=k2_relat_1(X14)|(~v1_relat_1(X14)|~v1_funct_1(X14)))&(X16=k1_funct_1(X14,esk5_3(X14,X15,X16))|~r2_hidden(X16,X15)|X15!=k2_relat_1(X14)|(~v1_relat_1(X14)|~v1_funct_1(X14))))&(~r2_hidden(X19,k1_relat_1(X14))|X18!=k1_funct_1(X14,X19)|r2_hidden(X18,X15)|X15!=k2_relat_1(X14)|(~v1_relat_1(X14)|~v1_funct_1(X14))))&((~r2_hidden(esk6_2(X14,X20),X20)|(~r2_hidden(X22,k1_relat_1(X14))|esk6_2(X14,X20)!=k1_funct_1(X14,X22))|X20=k2_relat_1(X14)|(~v1_relat_1(X14)|~v1_funct_1(X14)))&((r2_hidden(esk7_2(X14,X20),k1_relat_1(X14))|r2_hidden(esk6_2(X14,X20),X20)|X20=k2_relat_1(X14)|(~v1_relat_1(X14)|~v1_funct_1(X14)))&(esk6_2(X14,X20)=k1_funct_1(X14,esk7_2(X14,X20))|r2_hidden(esk6_2(X14,X20),X20)|X20=k2_relat_1(X14)|(~v1_relat_1(X14)|~v1_funct_1(X14)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_funct_1])])])])])])).
fof(c_0_13, plain, ![X42, X43, X44]:(~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(X42,X43)))|v1_relat_1(X44)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_14, negated_conjecture, ![X9]:((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,esk1_0,esk2_0)&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))&(((r2_hidden(esk4_1(X9),esk1_0)|~r2_hidden(X9,esk2_0))&(X9=k1_funct_1(esk3_0,esk4_1(X9))|~r2_hidden(X9,esk2_0)))&k2_relset_1(esk1_0,esk2_0,esk3_0)!=esk2_0)), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).
fof(c_0_15, plain, ![X30, X31, X32, X33, X34]:((~r1_tarski(X30,X31)|(~r2_hidden(X32,X30)|r2_hidden(X32,X31)))&((r2_hidden(esk8_2(X33,X34),X33)|r1_tarski(X33,X34))&(~r2_hidden(esk8_2(X33,X34),X34)|r1_tarski(X33,X34)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_16, plain, ![X40, X41]:((~v5_relat_1(X41,X40)|r1_tarski(k2_relat_1(X41),X40)|~v1_relat_1(X41))&(~r1_tarski(k2_relat_1(X41),X40)|v5_relat_1(X41,X40)|~v1_relat_1(X41))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d19_relat_1])])])).
fof(c_0_17, plain, ![X36, X37, X38]:((~v1_xboole_0(X36)|~r2_hidden(X37,X36))&(r2_hidden(esk9_1(X38),X38)|v1_xboole_0(X38))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_18, plain, (r2_hidden(X3,X4)|~r2_hidden(X1,k1_relat_1(X2))|X3!=k1_funct_1(X2,X1)|X4!=k2_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_21, plain, ![X24, X25, X26]:((((~v1_funct_2(X26,X24,X25)|X24=k1_relset_1(X24,X25,X26)|X25=k1_xboole_0|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(X24,X25))))&(X24!=k1_relset_1(X24,X25,X26)|v1_funct_2(X26,X24,X25)|X25=k1_xboole_0|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(X24,X25)))))&((~v1_funct_2(X26,X24,X25)|X24=k1_relset_1(X24,X25,X26)|X24!=k1_xboole_0|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(X24,X25))))&(X24!=k1_relset_1(X24,X25,X26)|v1_funct_2(X26,X24,X25)|X24!=k1_xboole_0|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(X24,X25))))))&((~v1_funct_2(X26,X24,X25)|X26=k1_xboole_0|X24=k1_xboole_0|X25!=k1_xboole_0|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(X24,X25))))&(X26!=k1_xboole_0|v1_funct_2(X26,X24,X25)|X24=k1_xboole_0|X25!=k1_xboole_0|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(X24,X25)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
fof(c_0_22, plain, ![X27, X28, X29]:(~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(X27,X28)))|k1_relset_1(X27,X28,X29)=k1_relat_1(X29)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_23, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, plain, (r1_tarski(k2_relat_1(X1),X2)|~v5_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relat_1(X1))|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))|~v1_funct_1(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_18])])).
cnf(c_0_27, plain, (r2_hidden(esk7_2(X1,X2),k1_relat_1(X1))|r2_hidden(esk6_2(X1,X2),X2)|X2=k2_relat_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_28, negated_conjecture, (v1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_29, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_30, plain, (X2=k1_relset_1(X2,X3,X1)|X3=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, negated_conjecture, (v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_32, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33, plain, (r2_hidden(X1,X2)|~v1_relat_1(X3)|~v5_relat_1(X3,X2)|~r2_hidden(X1,k2_relat_1(X3))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_34, plain, (esk6_2(X1,X2)=k1_funct_1(X1,esk7_2(X1,X2))|r2_hidden(esk6_2(X1,X2),X2)|X2=k2_relat_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_35, plain, ![X11, X12, X13]:((v4_relat_1(X13,X11)|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12))))&(v5_relat_1(X13,X12)|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_36, plain, ![X45, X46, X47]:(~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(X45,X46)))|k2_relset_1(X45,X46,X47)=k2_relat_1(X47)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
cnf(c_0_37, plain, (~v1_xboole_0(k2_relat_1(X1))|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_38, negated_conjecture, (X1=k2_relat_1(esk3_0)|r2_hidden(esk7_2(esk3_0,X1),k1_relat_1(esk3_0))|r2_hidden(esk6_2(esk3_0,X1),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])])).
cnf(c_0_39, negated_conjecture, (k1_relset_1(esk1_0,esk2_0,esk3_0)=esk1_0|esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_20]), c_0_31])])).
cnf(c_0_40, negated_conjecture, (k1_relset_1(esk1_0,esk2_0,esk3_0)=k1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_32, c_0_20])).
cnf(c_0_41, plain, (r2_hidden(esk5_3(X1,X2,X3),k1_relat_1(X1))|~r2_hidden(X3,X2)|X2!=k2_relat_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_42, negated_conjecture, (r2_hidden(esk4_1(X1),esk1_0)|~r2_hidden(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_43, plain, (r2_hidden(k1_funct_1(X1,X2),X3)|~v1_relat_1(X1)|~v5_relat_1(X1,X3)|~r2_hidden(X2,k1_relat_1(X1))|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_33, c_0_26])).
cnf(c_0_44, negated_conjecture, (k1_funct_1(esk3_0,esk7_2(esk3_0,X1))=esk6_2(esk3_0,X1)|X1=k2_relat_1(esk3_0)|r2_hidden(esk6_2(esk3_0,X1),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_28]), c_0_29])])).
cnf(c_0_45, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_46, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_47, negated_conjecture, (k1_relat_1(X1)=k2_relat_1(esk3_0)|r2_hidden(esk7_2(esk3_0,k1_relat_1(X1)),k1_relat_1(esk3_0))|~v1_xboole_0(k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_48, negated_conjecture, (k1_relat_1(esk3_0)=esk1_0|esk2_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_49, plain, (r2_hidden(esk5_3(X1,k2_relat_1(X1),X2),k1_relat_1(X1))|~v1_relat_1(X1)|~r2_hidden(X2,k2_relat_1(X1))|~v1_funct_1(X1)), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_50, negated_conjecture, (~v1_xboole_0(esk1_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_25, c_0_42])).
cnf(c_0_51, negated_conjecture, (X1=k2_relat_1(esk3_0)|r2_hidden(esk6_2(esk3_0,X1),X1)|r2_hidden(esk6_2(esk3_0,X1),X2)|~v5_relat_1(esk3_0,X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_28]), c_0_29])]), c_0_38])).
cnf(c_0_52, negated_conjecture, (v5_relat_1(esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_45, c_0_20])).
cnf(c_0_53, negated_conjecture, (k2_relset_1(esk1_0,esk2_0,esk3_0)!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_54, negated_conjecture, (k2_relset_1(esk1_0,esk2_0,esk3_0)=k2_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_46, c_0_20])).
cnf(c_0_55, negated_conjecture, (k2_relat_1(esk3_0)=esk1_0|esk2_0=k1_xboole_0|r2_hidden(esk7_2(esk3_0,esk1_0),k1_relat_1(esk3_0))|~v1_xboole_0(k2_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_28]), c_0_29])])).
cnf(c_0_56, plain, (~v1_xboole_0(k1_relat_1(X1))|~v1_relat_1(X1)|~r2_hidden(X2,k2_relat_1(X1))|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_25, c_0_49])).
cnf(c_0_57, plain, (r2_hidden(esk9_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_58, negated_conjecture, (~v1_xboole_0(esk1_0)|~v1_relat_1(X1)|~v5_relat_1(X1,esk2_0)|~r2_hidden(X2,k1_relat_1(X1))|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_50, c_0_43])).
cnf(c_0_59, negated_conjecture, (X1=k2_relat_1(esk3_0)|r2_hidden(esk6_2(esk3_0,X1),esk2_0)|r2_hidden(esk6_2(esk3_0,X1),X1)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_60, negated_conjecture, (k2_relat_1(esk3_0)!=esk2_0), inference(rw,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_61, negated_conjecture, (X1=k2_relat_1(esk3_0)|r2_hidden(esk6_2(esk3_0,X1),X1)|~v1_xboole_0(k1_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_25, c_0_38])).
cnf(c_0_62, negated_conjecture, (k2_relat_1(esk3_0)=esk1_0|esk2_0=k1_xboole_0|~v1_xboole_0(k2_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_55]), c_0_28]), c_0_29])])).
cnf(c_0_63, plain, (v1_xboole_0(k2_relat_1(X1))|~v1_xboole_0(k1_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_64, negated_conjecture, (X1=k2_relat_1(esk3_0)|r2_hidden(esk6_2(esk3_0,X1),X1)|~v1_xboole_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_38]), c_0_28]), c_0_52]), c_0_29])])).
cnf(c_0_65, plain, (v1_xboole_0(k2_relat_1(X1))|r2_hidden(esk9_1(k2_relat_1(X1)),X2)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(spm,[status(thm)],[c_0_33, c_0_57])).
cnf(c_0_66, plain, (X2=k2_relat_1(X1)|~r2_hidden(esk6_2(X1,X2),X2)|~r2_hidden(X3,k1_relat_1(X1))|esk6_2(X1,X2)!=k1_funct_1(X1,X3)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_67, negated_conjecture, (r2_hidden(esk6_2(esk3_0,esk2_0),esk2_0)), inference(sr,[status(thm)],[inference(ef,[status(thm)],[c_0_59]), c_0_60])).
cnf(c_0_68, negated_conjecture, (X1=k2_relat_1(esk3_0)|~v1_xboole_0(k1_relat_1(esk3_0))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_25, c_0_61])).
cnf(c_0_69, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_70, negated_conjecture, (k2_relat_1(esk3_0)=esk1_0|esk2_0=k1_xboole_0|~v1_xboole_0(k1_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_28]), c_0_29])])).
cnf(c_0_71, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_64]), c_0_60])).
cnf(c_0_72, plain, (v1_xboole_0(k2_relat_1(X1))|~v1_xboole_0(X2)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(spm,[status(thm)],[c_0_25, c_0_65])).
cnf(c_0_73, negated_conjecture, (esk6_2(esk3_0,esk2_0)!=k1_funct_1(esk3_0,X1)|~r2_hidden(X1,k1_relat_1(esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_28]), c_0_29])]), c_0_60])).
cnf(c_0_74, negated_conjecture, (X1=k1_funct_1(esk3_0,esk4_1(X1))|~r2_hidden(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_75, negated_conjecture, (k2_relat_1(esk3_0)=k1_xboole_0|~v1_xboole_0(k1_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_76, negated_conjecture, (esk2_0=k1_xboole_0|~v1_xboole_0(k1_relat_1(esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_70]), c_0_28]), c_0_29])]), c_0_71])).
cnf(c_0_77, plain, (v1_xboole_0(k1_relat_1(X1))|~v1_xboole_0(k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_37, c_0_57])).
cnf(c_0_78, negated_conjecture, (v1_xboole_0(k2_relat_1(esk3_0))|~v1_xboole_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_52]), c_0_28])])).
cnf(c_0_79, negated_conjecture, (~r2_hidden(esk4_1(esk6_2(esk3_0,esk2_0)),k1_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74])]), c_0_67])])).
cnf(c_0_80, negated_conjecture, (~v1_xboole_0(k1_relat_1(esk3_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_75]), c_0_76])).
cnf(c_0_81, negated_conjecture, (v1_xboole_0(k1_relat_1(esk3_0))|~v1_xboole_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_28]), c_0_29])])).
cnf(c_0_82, negated_conjecture, (esk2_0=k1_xboole_0|~r2_hidden(esk4_1(esk6_2(esk3_0,esk2_0)),esk1_0)), inference(spm,[status(thm)],[c_0_79, c_0_48])).
cnf(c_0_83, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_84, negated_conjecture, (esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_42]), c_0_67])])).
cnf(c_0_85, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_83, c_0_84]), c_0_69])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 86
# Proof object clause steps            : 64
# Proof object formula steps           : 22
# Proof object conjectures             : 42
# Proof object clause conjectures      : 39
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 39
# Proof object simplifying inferences  : 51
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 31
# Processed clauses                    : 510
# ...of these trivial                  : 1
# ...subsumed                          : 329
# ...remaining for further processing  : 180
# Other redundant clauses eliminated   : 15
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 22
# Backward-rewritten                   : 63
# Generated clauses                    : 803
# ...of the previous two non-trivial   : 773
# Contextual simplify-reflections      : 4
# Paramodulations                      : 786
# Factorizations                       : 4
# Equation resolutions                 : 15
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 88
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 80
# Current number of unprocessed clauses: 183
# ...number of literals in the above   : 1056
# Current number of archived formulas  : 0
# Current number of archived clauses   : 85
# Clause-clause subsumption calls (NU) : 4812
# Rec. Clause-clause subsumption calls : 1758
# Non-unit clause-clause subsumptions  : 161
# Unit Clause-clause subsumption calls : 172
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 14134

# -------------------------------------------------
# User time                : 0.021 s
# System time              : 0.000 s
# Total time               : 0.021 s
# Maximum resident set size: 28624 pages
