# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/4f7f0f8039214de2a6a7b80b1fe1bb07.E_input', d2_tarski)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/4f7f0f8039214de2a6a7b80b1fe1bb07.E_input', d9_lattice3)).
fof(t8_yellow_0, conjecture, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4)))&(((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4))=>r1_lattice3(X1,k2_tarski(X3,X4),X2))&((r2_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2)))&((r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2))=>r2_lattice3(X1,k2_tarski(X3,X4),X2))))))))), file('data/MPTP/proofs/4f7f0f8039214de2a6a7b80b1fe1bb07.E_input', t8_yellow_0)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/4f7f0f8039214de2a6a7b80b1fe1bb07.E_input', d8_lattice3)).
fof(c_0_4, plain, ![X9, X10, X11, X12, X13, X14, X15, X16]:(((~r2_hidden(X12,X11)|(X12=X9|X12=X10)|X11!=k2_tarski(X9,X10))&((X13!=X9|r2_hidden(X13,X11)|X11!=k2_tarski(X9,X10))&(X13!=X10|r2_hidden(X13,X11)|X11!=k2_tarski(X9,X10))))&(((esk5_3(X14,X15,X16)!=X14|~r2_hidden(esk5_3(X14,X15,X16),X16)|X16=k2_tarski(X14,X15))&(esk5_3(X14,X15,X16)!=X15|~r2_hidden(esk5_3(X14,X15,X16),X16)|X16=k2_tarski(X14,X15)))&(r2_hidden(esk5_3(X14,X15,X16),X16)|(esk5_3(X14,X15,X16)=X14|esk5_3(X14,X15,X16)=X15)|X16=k2_tarski(X14,X15)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_5, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_6, plain, ![X23, X24, X25, X26]:((~r2_lattice3(X23,X24,X25)|(~m1_subset_1(X26,u1_struct_0(X23))|(~r2_hidden(X26,X24)|r1_orders_2(X23,X26,X25)))|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))&((m1_subset_1(esk7_3(X23,X24,X25),u1_struct_0(X23))|r2_lattice3(X23,X24,X25)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))&((r2_hidden(esk7_3(X23,X24,X25),X24)|r2_lattice3(X23,X24,X25)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))&(~r1_orders_2(X23,esk7_3(X23,X24,X25),X25)|r2_lattice3(X23,X24,X25)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
fof(c_0_7, negated_conjecture, ~(![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4)))&(((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X2,X4))=>r1_lattice3(X1,k2_tarski(X3,X4),X2))&((r2_lattice3(X1,k2_tarski(X3,X4),X2)=>(r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2)))&((r1_orders_2(X1,X3,X2)&r1_orders_2(X1,X4,X2))=>r2_lattice3(X1,k2_tarski(X3,X4),X2)))))))))), inference(assume_negation,[status(cth)],[t8_yellow_0])).
cnf(c_0_8, plain, (X1=X2|X1=X3|~r2_hidden(X1,k2_tarski(X2,X3))), inference(er,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (r2_hidden(esk7_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_10, negated_conjecture, (l1_orders_2(esk1_0)&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(m1_subset_1(esk4_0,u1_struct_0(esk1_0))&(((((((r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_orders_2(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0))&(r1_orders_2(esk1_0,esk4_0,esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_orders_2(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)))&(~r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_orders_2(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)))&(((r1_orders_2(esk1_0,esk3_0,esk2_0)|(~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0))|r1_orders_2(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0))&(r1_orders_2(esk1_0,esk4_0,esk2_0)|(~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0))|r1_orders_2(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)))&(~r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|(~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0))|r1_orders_2(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0))))&((((r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_orders_2(esk1_0,esk2_0,esk4_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0))&(r1_orders_2(esk1_0,esk4_0,esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_orders_2(esk1_0,esk2_0,esk4_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)))&(~r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_orders_2(esk1_0,esk2_0,esk4_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)))&(((r1_orders_2(esk1_0,esk3_0,esk2_0)|(~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0))|r1_orders_2(esk1_0,esk2_0,esk4_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0))&(r1_orders_2(esk1_0,esk4_0,esk2_0)|(~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0))|r1_orders_2(esk1_0,esk2_0,esk4_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)))&(~r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|(~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0))|r1_orders_2(esk1_0,esk2_0,esk4_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)))))&((((r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0))&(r1_orders_2(esk1_0,esk4_0,esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)))&(~r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)))&(((r1_orders_2(esk1_0,esk3_0,esk2_0)|(~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0))|~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0))&(r1_orders_2(esk1_0,esk4_0,esk2_0)|(~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0))|~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)))&(~r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|(~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0))|~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)))))&((((((r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_orders_2(esk1_0,esk2_0,esk3_0)|(~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0)))&(r1_orders_2(esk1_0,esk4_0,esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_orders_2(esk1_0,esk2_0,esk3_0)|(~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0))))&(~r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_orders_2(esk1_0,esk2_0,esk3_0)|(~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0))))&(((r1_orders_2(esk1_0,esk3_0,esk2_0)|(~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0))|r1_orders_2(esk1_0,esk2_0,esk3_0)|(~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0)))&(r1_orders_2(esk1_0,esk4_0,esk2_0)|(~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0))|r1_orders_2(esk1_0,esk2_0,esk3_0)|(~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0))))&(~r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|(~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0))|r1_orders_2(esk1_0,esk2_0,esk3_0)|(~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0)))))&((((r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_orders_2(esk1_0,esk2_0,esk4_0)|(~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0)))&(r1_orders_2(esk1_0,esk4_0,esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_orders_2(esk1_0,esk2_0,esk4_0)|(~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0))))&(~r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_orders_2(esk1_0,esk2_0,esk4_0)|(~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0))))&(((r1_orders_2(esk1_0,esk3_0,esk2_0)|(~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0))|r1_orders_2(esk1_0,esk2_0,esk4_0)|(~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0)))&(r1_orders_2(esk1_0,esk4_0,esk2_0)|(~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0))|r1_orders_2(esk1_0,esk2_0,esk4_0)|(~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0))))&(~r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|(~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0))|r1_orders_2(esk1_0,esk2_0,esk4_0)|(~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0))))))&((((r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|(~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0)))&(r1_orders_2(esk1_0,esk4_0,esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|(~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0))))&(~r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|(~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0))))&(((r1_orders_2(esk1_0,esk3_0,esk2_0)|(~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0))|~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|(~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0)))&(r1_orders_2(esk1_0,esk4_0,esk2_0)|(~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0))|~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|(~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0))))&(~r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|(~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0))|~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|(~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0))))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])).
cnf(c_0_11, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_12, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_13, plain, (esk7_3(X1,k2_tarski(X2,X3),X4)=X3|esk7_3(X1,k2_tarski(X2,X3),X4)=X2|r2_lattice3(X1,k2_tarski(X2,X3),X4)|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_16, plain, ![X18, X19, X20, X21]:((~r1_lattice3(X18,X19,X20)|(~m1_subset_1(X21,u1_struct_0(X18))|(~r2_hidden(X21,X19)|r1_orders_2(X18,X20,X21)))|~m1_subset_1(X20,u1_struct_0(X18))|~l1_orders_2(X18))&((m1_subset_1(esk6_3(X18,X19,X20),u1_struct_0(X18))|r1_lattice3(X18,X19,X20)|~m1_subset_1(X20,u1_struct_0(X18))|~l1_orders_2(X18))&((r2_hidden(esk6_3(X18,X19,X20),X19)|r1_lattice3(X18,X19,X20)|~m1_subset_1(X20,u1_struct_0(X18))|~l1_orders_2(X18))&(~r1_orders_2(X18,X20,esk6_3(X18,X19,X20))|r1_lattice3(X18,X19,X20)|~m1_subset_1(X20,u1_struct_0(X18))|~l1_orders_2(X18))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
cnf(c_0_17, plain, (r1_orders_2(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_18, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X3,X1)), inference(er,[status(thm)],[c_0_11])).
cnf(c_0_20, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_21, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X1,X3)), inference(er,[status(thm)],[c_0_12])).
cnf(c_0_22, plain, (r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk7_3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_23, negated_conjecture, (esk7_3(esk1_0,k2_tarski(X1,X2),esk2_0)=X1|esk7_3(esk1_0,k2_tarski(X1,X2),esk2_0)=X2|r2_lattice3(esk1_0,k2_tarski(X1,X2),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])])).
cnf(c_0_24, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,X1)|~r2_hidden(esk4_0,X2)|~r2_lattice3(esk1_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_15])])).
cnf(c_0_26, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_orders_2(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_27, plain, (r2_hidden(X1,k2_tarski(X2,X1))), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_28, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,X1)|~r2_hidden(esk3_0,X2)|~r2_lattice3(esk1_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_20]), c_0_15])])).
cnf(c_0_29, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_orders_2(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_30, plain, (r2_hidden(X1,k2_tarski(X1,X2))), inference(er,[status(thm)],[c_0_21])).
cnf(c_0_31, negated_conjecture, (esk7_3(esk1_0,k2_tarski(X1,X2),esk2_0)=X1|r2_lattice3(esk1_0,k2_tarski(X1,X2),esk2_0)|~r1_orders_2(esk1_0,X2,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_14]), c_0_15])])).
cnf(c_0_32, negated_conjecture, (r1_orders_2(esk1_0,X1,esk3_0)|~r2_hidden(esk3_0,X2)|~r1_lattice3(esk1_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_20]), c_0_15])])).
cnf(c_0_33, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)|r1_orders_2(esk1_0,esk4_0,esk2_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_14])])).
cnf(c_0_34, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)|r1_orders_2(esk1_0,esk3_0,esk2_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30]), c_0_14])])).
cnf(c_0_35, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_orders_2(esk1_0,esk2_0,esk4_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_36, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|r1_orders_2(esk1_0,esk2_0,esk4_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_37, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_38, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_39, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|~r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_40, negated_conjecture, (r2_lattice3(esk1_0,k2_tarski(X1,X2),esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,X2,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_31]), c_0_14]), c_0_15])])).
cnf(c_0_41, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,esk2_0)|r1_orders_2(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_30]), c_0_14])])).
cnf(c_0_42, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk2_0)|r1_orders_2(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_34]), c_0_30]), c_0_14])])).
cnf(c_0_43, negated_conjecture, (r1_orders_2(esk1_0,X1,esk4_0)|~r2_hidden(esk4_0,X2)|~r1_lattice3(esk1_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_18]), c_0_15])])).
cnf(c_0_44, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk4_0)|r1_orders_2(esk1_0,esk4_0,esk2_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_35]), c_0_27]), c_0_14])])).
cnf(c_0_45, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk4_0)|r1_orders_2(esk1_0,esk3_0,esk2_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_36]), c_0_30]), c_0_14])])).
cnf(c_0_46, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_47, negated_conjecture, (~r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0)|~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_48, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,esk2_0)|~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0)|~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_37]), c_0_27]), c_0_14])])).
cnf(c_0_49, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0)|~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_38]), c_0_30]), c_0_14])])).
cnf(c_0_50, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41]), c_0_42])).
cnf(c_0_51, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk4_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|~r2_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)|~r1_orders_2(esk1_0,esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_52, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,esk2_0)|r1_orders_2(esk1_0,esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_27]), c_0_14])])).
cnf(c_0_53, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk2_0)|r1_orders_2(esk1_0,esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_45]), c_0_27]), c_0_14])])).
cnf(c_0_54, plain, (esk6_3(X1,k2_tarski(X2,X3),X4)=X3|esk6_3(X1,k2_tarski(X2,X3),X4)=X2|r1_lattice3(X1,k2_tarski(X2,X3),X4)|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_8, c_0_46])).
cnf(c_0_55, negated_conjecture, (~r1_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk2_0,esk4_0)|~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_40]), c_0_48]), c_0_49])).
cnf(c_0_56, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_50]), c_0_30]), c_0_14])])).
cnf(c_0_57, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk4_0)|r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_40]), c_0_52]), c_0_53])).
cnf(c_0_58, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk6_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_59, negated_conjecture, (esk6_3(esk1_0,k2_tarski(X1,X2),esk2_0)=X1|esk6_3(esk1_0,k2_tarski(X1,X2),esk2_0)=X2|r1_lattice3(esk1_0,k2_tarski(X1,X2),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_14]), c_0_15])])).
cnf(c_0_60, negated_conjecture, (~r1_orders_2(esk1_0,esk2_0,esk4_0)|~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_56])])).
cnf(c_0_61, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_57]), c_0_27]), c_0_14])])).
cnf(c_0_62, negated_conjecture, (esk6_3(esk1_0,k2_tarski(X1,X2),esk2_0)=X1|r1_lattice3(esk1_0,k2_tarski(X1,X2),esk2_0)|~r1_orders_2(esk1_0,esk2_0,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_14]), c_0_15])])).
cnf(c_0_63, negated_conjecture, (~r1_lattice3(esk1_0,k2_tarski(esk3_0,esk4_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_61])])).
cnf(c_0_64, negated_conjecture, (r1_lattice3(esk1_0,k2_tarski(X1,X2),esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,esk2_0,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_62]), c_0_14]), c_0_15])])).
cnf(c_0_65, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_56]), c_0_61])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 66
# Proof object clause steps            : 57
# Proof object formula steps           : 9
# Proof object conjectures             : 44
# Proof object clause conjectures      : 41
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 31
# Proof object simplifying inferences  : 75
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 54
# Removed in clause preprocessing      : 27
# Initial clauses in saturation        : 27
# Processed clauses                    : 910
# ...of these trivial                  : 0
# ...subsumed                          : 460
# ...remaining for further processing  : 450
# Other redundant clauses eliminated   : 92
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 19
# Backward-rewritten                   : 20
# Generated clauses                    : 3733
# ...of the previous two non-trivial   : 2668
# Contextual simplify-reflections      : 27
# Paramodulations                      : 3578
# Factorizations                       : 60
# Equation resolutions                 : 95
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 382
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 373
# Current number of unprocessed clauses: 1773
# ...number of literals in the above   : 10046
# Current number of archived formulas  : 0
# Current number of archived clauses   : 66
# Clause-clause subsumption calls (NU) : 21018
# Rec. Clause-clause subsumption calls : 4950
# Non-unit clause-clause subsumptions  : 506
# Unit Clause-clause subsumption calls : 23
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 95680

# -------------------------------------------------
# User time                : 0.048 s
# System time              : 0.003 s
# Total time               : 0.051 s
# Maximum resident set size: 28764 pages
