# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/50bb7b0c5b85480cb9ab7830bb59bcb7.E_input', d1_xboole_0)).
fof(t103_zfmisc_1, axiom, ![X1, X2, X3, X4]:~((r1_tarski(X1,k2_zfmisc_1(X2,X3))&(r2_hidden(X4,X1)&![X5, X6]:~((r2_hidden(X5,X2)&(r2_hidden(X6,X3)&X4=k4_tarski(X5,X6))))))), file('data/MPTP/proofs/50bb7b0c5b85480cb9ab7830bb59bcb7.E_input', t103_zfmisc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/50bb7b0c5b85480cb9ab7830bb59bcb7.E_input', d10_xboole_0)).
fof(t69_mcart_1, conjecture, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X6))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k5_mcart_1(X1,X2,X3,X5)))))), file('data/MPTP/proofs/50bb7b0c5b85480cb9ab7830bb59bcb7.E_input', t69_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/50bb7b0c5b85480cb9ab7830bb59bcb7.E_input', d3_zfmisc_1)).
fof(t50_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))&(k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))&k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4))))))))), file('data/MPTP/proofs/50bb7b0c5b85480cb9ab7830bb59bcb7.E_input', t50_mcart_1)).
fof(d4_zfmisc_1, axiom, ![X1, X2, X3, X4]:k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4), file('data/MPTP/proofs/50bb7b0c5b85480cb9ab7830bb59bcb7.E_input', d4_zfmisc_1)).
fof(t34_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4, X5, X6]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k4_mcart_1(X3,X4,X5,X6))))))), file('data/MPTP/proofs/50bb7b0c5b85480cb9ab7830bb59bcb7.E_input', t34_mcart_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/50bb7b0c5b85480cb9ab7830bb59bcb7.E_input', t2_subset)).
fof(t55_mcart_1, axiom, ![X1, X2, X3, X4]:((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(X4=k1_xboole_0))))<=>~(k4_zfmisc_1(X1,X2,X3,X4)=k1_xboole_0)), file('data/MPTP/proofs/50bb7b0c5b85480cb9ab7830bb59bcb7.E_input', t55_mcart_1)).
fof(t10_mcart_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(r2_hidden(k1_mcart_1(X1),X2)&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/50bb7b0c5b85480cb9ab7830bb59bcb7.E_input', t10_mcart_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/50bb7b0c5b85480cb9ab7830bb59bcb7.E_input', d3_mcart_1)).
fof(t23_mcart_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r2_hidden(X1,X2)=>X1=k4_tarski(k1_mcart_1(X1),k2_mcart_1(X1)))), file('data/MPTP/proofs/50bb7b0c5b85480cb9ab7830bb59bcb7.E_input', t23_mcart_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/50bb7b0c5b85480cb9ab7830bb59bcb7.E_input', fc6_relat_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/50bb7b0c5b85480cb9ab7830bb59bcb7.E_input', t1_subset)).
fof(dt_k7_mcart_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>m1_subset_1(k7_mcart_1(X1,X2,X3,X4),X3)), file('data/MPTP/proofs/50bb7b0c5b85480cb9ab7830bb59bcb7.E_input', dt_k7_mcart_1)).
fof(c_0_16, plain, ![X9, X10, X11]:((~v1_xboole_0(X9)|~r2_hidden(X10,X9))&(r2_hidden(esk1_1(X11),X11)|v1_xboole_0(X11))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_17, plain, ![X13, X14, X15, X16]:((r2_hidden(esk2_4(X13,X14,X15,X16),X14)|~r2_hidden(X16,X13)|~r1_tarski(X13,k2_zfmisc_1(X14,X15)))&((r2_hidden(esk3_4(X13,X14,X15,X16),X15)|~r2_hidden(X16,X13)|~r1_tarski(X13,k2_zfmisc_1(X14,X15)))&(X16=k4_tarski(esk2_4(X13,X14,X15,X16),esk3_4(X13,X14,X15,X16))|~r2_hidden(X16,X13)|~r1_tarski(X13,k2_zfmisc_1(X14,X15))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t103_zfmisc_1])])])])])])).
fof(c_0_18, plain, ![X19, X20]:(((r1_tarski(X19,X20)|X19!=X20)&(r1_tarski(X20,X19)|X19!=X20))&(~r1_tarski(X19,X20)|~r1_tarski(X20,X19)|X19=X20)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_19, negated_conjecture, ~(![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X6))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k5_mcart_1(X1,X2,X3,X5))))))), inference(assume_negation,[status(cth)],[t69_mcart_1])).
cnf(c_0_20, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, plain, (r2_hidden(esk2_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X1)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_23, negated_conjecture, ![X26, X27, X28]:(m1_subset_1(esk8_0,k3_zfmisc_1(esk4_0,esk5_0,esk6_0))&((~m1_subset_1(X26,esk4_0)|(~m1_subset_1(X27,esk5_0)|(~m1_subset_1(X28,esk6_0)|(esk8_0!=k3_mcart_1(X26,X27,X28)|esk7_0=X26))))&(esk4_0!=k1_xboole_0&(esk5_0!=k1_xboole_0&(esk6_0!=k1_xboole_0&esk7_0!=k5_mcart_1(esk4_0,esk5_0,esk6_0,esk8_0)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])).
fof(c_0_24, plain, ![X29, X30, X31]:k3_zfmisc_1(X29,X30,X31)=k2_zfmisc_1(k2_zfmisc_1(X29,X30),X31), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_25, plain, ![X32, X33, X34, X35]:((k5_mcart_1(X32,X33,X34,X35)=k1_mcart_1(k1_mcart_1(X35))|~m1_subset_1(X35,k3_zfmisc_1(X32,X33,X34))|X34=k1_xboole_0|X33=k1_xboole_0|X32=k1_xboole_0)&((k6_mcart_1(X32,X33,X34,X35)=k2_mcart_1(k1_mcart_1(X35))|~m1_subset_1(X35,k3_zfmisc_1(X32,X33,X34))|X34=k1_xboole_0|X33=k1_xboole_0|X32=k1_xboole_0)&(k7_mcart_1(X32,X33,X34,X35)=k2_mcart_1(X35)|~m1_subset_1(X35,k3_zfmisc_1(X32,X33,X34))|X34=k1_xboole_0|X33=k1_xboole_0|X32=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_mcart_1])])])])])).
fof(c_0_26, plain, ![X36, X37, X38, X39]:k4_zfmisc_1(X36,X37,X38,X39)=k2_zfmisc_1(k3_zfmisc_1(X36,X37,X38),X39), inference(variable_rename,[status(thm)],[d4_zfmisc_1])).
cnf(c_0_27, plain, (~r1_tarski(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X4,X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_28, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_22])).
fof(c_0_29, plain, ![X40, X42, X43, X44, X45]:((r2_hidden(esk9_1(X40),X40)|X40=k1_xboole_0)&((~r2_hidden(X42,X40)|esk9_1(X40)!=k4_mcart_1(X42,X43,X44,X45)|X40=k1_xboole_0)&(~r2_hidden(X43,X40)|esk9_1(X40)!=k4_mcart_1(X42,X43,X44,X45)|X40=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_mcart_1])])])])])])])).
fof(c_0_30, plain, ![X46, X47]:(~m1_subset_1(X46,X47)|(v1_xboole_0(X47)|r2_hidden(X46,X47))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk8_0,k3_zfmisc_1(esk4_0,esk5_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, plain, (k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_34, plain, ![X48, X49, X50, X51]:((X48=k1_xboole_0|(X49=k1_xboole_0|(X50=k1_xboole_0|X51=k1_xboole_0))|k4_zfmisc_1(X48,X49,X50,X51)!=k1_xboole_0)&((X48!=k1_xboole_0|k4_zfmisc_1(X48,X49,X50,X51)=k1_xboole_0)&((X49!=k1_xboole_0|k4_zfmisc_1(X48,X49,X50,X51)=k1_xboole_0)&((X50!=k1_xboole_0|k4_zfmisc_1(X48,X49,X50,X51)=k1_xboole_0)&(X51!=k1_xboole_0|k4_zfmisc_1(X48,X49,X50,X51)=k1_xboole_0))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t55_mcart_1])])])])).
cnf(c_0_35, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (~r2_hidden(X1,k2_zfmisc_1(X2,X3))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_37, plain, (r2_hidden(esk9_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39, negated_conjecture, (m1_subset_1(esk8_0,k2_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0),esk6_0))), inference(rw,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_40, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_33, c_0_32])).
cnf(c_0_41, negated_conjecture, (esk6_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_42, negated_conjecture, (esk5_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_43, negated_conjecture, (esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_44, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|k4_zfmisc_1(X1,X2,X3,X4)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_35, c_0_32])).
cnf(c_0_46, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_47, negated_conjecture, (r2_hidden(esk8_0,k2_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0),esk6_0))|v1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0),esk6_0))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_48, negated_conjecture, (esk7_0!=k5_mcart_1(esk4_0,esk5_0,esk6_0,esk8_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_49, negated_conjecture, (k5_mcart_1(esk4_0,esk5_0,esk6_0,esk8_0)=k1_mcart_1(k1_mcart_1(esk8_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_39]), c_0_41]), c_0_42]), c_0_43])).
cnf(c_0_50, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_51, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0),esk6_0),X1)=k1_xboole_0|r2_hidden(esk8_0,k2_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0),esk6_0))), inference(spm,[status(thm)],[c_0_46, c_0_47])).
fof(c_0_52, plain, ![X52, X53, X54]:((r2_hidden(k1_mcart_1(X52),X53)|~r2_hidden(X52,k2_zfmisc_1(X53,X54)))&(r2_hidden(k2_mcart_1(X52),X54)|~r2_hidden(X52,k2_zfmisc_1(X53,X54)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_mcart_1])])])).
cnf(c_0_53, negated_conjecture, (k1_mcart_1(k1_mcart_1(esk8_0))!=esk7_0), inference(rw,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_54, negated_conjecture, (X1=k1_xboole_0|r2_hidden(esk8_0,k2_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0),esk6_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_41]), c_0_42]), c_0_43])).
fof(c_0_55, plain, ![X55, X56, X57]:k3_mcart_1(X55,X56,X57)=k4_tarski(k4_tarski(X55,X56),X57), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_56, plain, ![X58, X59]:(~v1_relat_1(X59)|(~r2_hidden(X58,X59)|X58=k4_tarski(k1_mcart_1(X58),k2_mcart_1(X58)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_mcart_1])])).
cnf(c_0_57, plain, (r2_hidden(k1_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_58, negated_conjecture, (r2_hidden(esk8_0,k2_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0),esk6_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_54])).
fof(c_0_59, plain, ![X60, X61]:v1_relat_1(k2_zfmisc_1(X60,X61)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_60, negated_conjecture, (esk7_0=X1|~m1_subset_1(X1,esk4_0)|~m1_subset_1(X2,esk5_0)|~m1_subset_1(X3,esk6_0)|esk8_0!=k3_mcart_1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_61, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_62, plain, (X2=k4_tarski(k1_mcart_1(X2),k2_mcart_1(X2))|~v1_relat_1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_63, negated_conjecture, (r2_hidden(k1_mcart_1(esk8_0),k2_zfmisc_1(esk4_0,esk5_0))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_64, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_65, negated_conjecture, (esk7_0=X1|esk8_0!=k4_tarski(k4_tarski(X1,X2),X3)|~m1_subset_1(X3,esk6_0)|~m1_subset_1(X2,esk5_0)|~m1_subset_1(X1,esk4_0)), inference(rw,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_66, negated_conjecture, (k4_tarski(k1_mcart_1(k1_mcart_1(esk8_0)),k2_mcart_1(k1_mcart_1(esk8_0)))=k1_mcart_1(esk8_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64])])).
fof(c_0_67, plain, ![X62, X63]:(~r2_hidden(X62,X63)|m1_subset_1(X62,X63)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_68, plain, (r2_hidden(k2_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_52])).
fof(c_0_69, plain, ![X64, X65, X66, X67]:(~m1_subset_1(X67,k3_zfmisc_1(X64,X65,X66))|m1_subset_1(k7_mcart_1(X64,X65,X66,X67),X66)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_mcart_1])])).
cnf(c_0_70, plain, (k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_71, negated_conjecture, (k4_tarski(k1_mcart_1(esk8_0),X1)!=esk8_0|~m1_subset_1(k2_mcart_1(k1_mcart_1(esk8_0)),esk5_0)|~m1_subset_1(k1_mcart_1(k1_mcart_1(esk8_0)),esk4_0)|~m1_subset_1(X1,esk6_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_53])).
cnf(c_0_72, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_73, negated_conjecture, (r2_hidden(k2_mcart_1(k1_mcart_1(esk8_0)),esk5_0)), inference(spm,[status(thm)],[c_0_68, c_0_63])).
cnf(c_0_74, plain, (m1_subset_1(k7_mcart_1(X2,X3,X4,X1),X4)|~m1_subset_1(X1,k3_zfmisc_1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_75, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_70, c_0_32])).
cnf(c_0_76, negated_conjecture, (k4_tarski(k1_mcart_1(esk8_0),X1)!=esk8_0|~m1_subset_1(k1_mcart_1(k1_mcart_1(esk8_0)),esk4_0)|~m1_subset_1(X1,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73])])).
cnf(c_0_77, negated_conjecture, (r2_hidden(k1_mcart_1(k1_mcart_1(esk8_0)),esk4_0)), inference(spm,[status(thm)],[c_0_57, c_0_63])).
cnf(c_0_78, plain, (m1_subset_1(k7_mcart_1(X2,X3,X4,X1),X4)|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4))), inference(rw,[status(thm)],[c_0_74, c_0_32])).
cnf(c_0_79, negated_conjecture, (k7_mcart_1(esk4_0,esk5_0,esk6_0,esk8_0)=k2_mcart_1(esk8_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_39]), c_0_41]), c_0_42]), c_0_43])).
cnf(c_0_80, negated_conjecture, (k4_tarski(k1_mcart_1(esk8_0),X1)!=esk8_0|~m1_subset_1(X1,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_72]), c_0_77])])).
cnf(c_0_81, negated_conjecture, (k4_tarski(k1_mcart_1(esk8_0),k2_mcart_1(esk8_0))=esk8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_58]), c_0_64])])).
cnf(c_0_82, negated_conjecture, (m1_subset_1(k2_mcart_1(esk8_0),esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_39])])).
cnf(c_0_83, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_82])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 84
# Proof object clause steps            : 51
# Proof object formula steps           : 33
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 32
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 32
# Processed clauses                    : 316
# ...of these trivial                  : 10
# ...subsumed                          : 155
# ...remaining for further processing  : 151
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 25
# Backward-rewritten                   : 5
# Generated clauses                    : 667
# ...of the previous two non-trivial   : 525
# Contextual simplify-reflections      : 1
# Paramodulations                      : 657
# Factorizations                       : 0
# Equation resolutions                 : 9
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 87
#    Positive orientable unit clauses  : 16
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 62
# Current number of unprocessed clauses: 229
# ...number of literals in the above   : 903
# Current number of archived formulas  : 0
# Current number of archived clauses   : 65
# Clause-clause subsumption calls (NU) : 904
# Rec. Clause-clause subsumption calls : 758
# Non-unit clause-clause subsumptions  : 167
# Unit Clause-clause subsumption calls : 89
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8938

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.003 s
# Total time               : 0.014 s
# Maximum resident set size: 28684 pages
