# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d20_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v13_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r1_orders_2(X1,X3,X4))=>r2_hidden(X4,X2))))))), file('data/MPTP/proofs/514981734f954d89a18b14b3ac11c99a.E_input', d20_waybel_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/514981734f954d89a18b14b3ac11c99a.E_input', t4_subset)).
fof(t8_waybel_7, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v1_subset_1(X2,u1_struct_0(X1))<=>~(r2_hidden(k3_yellow_0(X1),X2))))), file('data/MPTP/proofs/514981734f954d89a18b14b3ac11c99a.E_input', t8_waybel_7)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/514981734f954d89a18b14b3ac11c99a.E_input', t1_subset)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/514981734f954d89a18b14b3ac11c99a.E_input', d3_tarski)).
fof(t44_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,k3_yellow_0(X1),X2))), file('data/MPTP/proofs/514981734f954d89a18b14b3ac11c99a.E_input', t44_yellow_0)).
fof(d7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(v1_subset_1(X2,X1)<=>~(X2=X1))), file('data/MPTP/proofs/514981734f954d89a18b14b3ac11c99a.E_input', d7_subset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/514981734f954d89a18b14b3ac11c99a.E_input', t3_subset)).
fof(t8_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,X2)<=>r2_hidden(X4,X3)))=>X2=X3))), file('data/MPTP/proofs/514981734f954d89a18b14b3ac11c99a.E_input', t8_subset_1)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/514981734f954d89a18b14b3ac11c99a.E_input', l3_subset_1)).
fof(dt_k3_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(k3_yellow_0(X1),u1_struct_0(X1))), file('data/MPTP/proofs/514981734f954d89a18b14b3ac11c99a.E_input', dt_k3_yellow_0)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/514981734f954d89a18b14b3ac11c99a.E_input', t2_subset)).
fof(c_0_12, plain, ![X5, X6, X7, X8]:((~v13_waybel_0(X6,X5)|(~m1_subset_1(X7,u1_struct_0(X5))|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X7,X6)|~r1_orders_2(X5,X7,X8)|r2_hidden(X8,X6))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))|~l1_orders_2(X5))&((m1_subset_1(esk1_2(X5,X6),u1_struct_0(X5))|v13_waybel_0(X6,X5)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))|~l1_orders_2(X5))&((m1_subset_1(esk2_2(X5,X6),u1_struct_0(X5))|v13_waybel_0(X6,X5)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))|~l1_orders_2(X5))&(((r2_hidden(esk1_2(X5,X6),X6)|v13_waybel_0(X6,X5)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))|~l1_orders_2(X5))&(r1_orders_2(X5,esk1_2(X5,X6),esk2_2(X5,X6))|v13_waybel_0(X6,X5)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))|~l1_orders_2(X5)))&(~r2_hidden(esk2_2(X5,X6),X6)|v13_waybel_0(X6,X5)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d20_waybel_0])])])])])).
fof(c_0_13, plain, ![X11, X12, X13]:(~r2_hidden(X11,X12)|~m1_subset_1(X12,k1_zfmisc_1(X13))|m1_subset_1(X11,X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_14, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v1_subset_1(X2,u1_struct_0(X1))<=>~(r2_hidden(k3_yellow_0(X1),X2)))))), inference(assume_negation,[status(cth)],[t8_waybel_7])).
fof(c_0_15, plain, ![X16, X17]:(~r2_hidden(X16,X17)|m1_subset_1(X16,X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
fof(c_0_16, plain, ![X18, X19, X20, X21, X22]:((~r1_tarski(X18,X19)|(~r2_hidden(X20,X18)|r2_hidden(X20,X19)))&((r2_hidden(esk5_2(X21,X22),X21)|r1_tarski(X21,X22))&(~r2_hidden(esk5_2(X21,X22),X22)|r1_tarski(X21,X22)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_17, plain, (r2_hidden(X4,X1)|~v13_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,X1)|~r1_orders_2(X2,X3,X4)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_19, negated_conjecture, ((~v2_struct_0(esk3_0)&(v3_orders_2(esk3_0)&(v4_orders_2(esk3_0)&(v5_orders_2(esk3_0)&(v1_yellow_0(esk3_0)&l1_orders_2(esk3_0))))))&((~v1_xboole_0(esk4_0)&(v2_waybel_0(esk4_0,esk3_0)&(v13_waybel_0(esk4_0,esk3_0)&m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk3_0))))))&((~v1_subset_1(esk4_0,u1_struct_0(esk3_0))|r2_hidden(k3_yellow_0(esk3_0),esk4_0))&(v1_subset_1(esk4_0,u1_struct_0(esk3_0))|~r2_hidden(k3_yellow_0(esk3_0),esk4_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])).
fof(c_0_20, plain, ![X24, X25]:(v2_struct_0(X24)|(~v5_orders_2(X24)|(~v1_yellow_0(X24)|~l1_orders_2(X24)))|(~m1_subset_1(X25,u1_struct_0(X24))|r1_orders_2(X24,k3_yellow_0(X24),X25))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t44_yellow_0])])])])).
cnf(c_0_21, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (r2_hidden(X1,X2)|~r1_orders_2(X3,X4,X1)|~r2_hidden(X4,X2)|~v13_waybel_0(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)), inference(csr,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_24, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, negated_conjecture, (v13_waybel_0(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, negated_conjecture, (l1_orders_2(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (v2_struct_0(X1)|r1_orders_2(X1,k3_yellow_0(X1),X2)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (r1_tarski(X1,X2)|m1_subset_1(esk5_2(X1,X2),X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
fof(c_0_29, plain, ![X26, X27]:((~v1_subset_1(X27,X26)|X27!=X26|~m1_subset_1(X27,k1_zfmisc_1(X26)))&(X27=X26|v1_subset_1(X27,X26)|~m1_subset_1(X27,k1_zfmisc_1(X26)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_subset_1])])])])).
cnf(c_0_30, negated_conjecture, (r2_hidden(X1,esk4_0)|~r1_orders_2(esk3_0,X2,X1)|~r2_hidden(X2,esk4_0)|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25]), c_0_26])])).
cnf(c_0_31, plain, (r1_tarski(u1_struct_0(X1),X2)|v2_struct_0(X1)|r1_orders_2(X1,k3_yellow_0(X1),esk5_2(u1_struct_0(X1),X2))|~v1_yellow_0(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_32, negated_conjecture, (v1_yellow_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_33, negated_conjecture, (v5_orders_2(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_34, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_35, plain, (X1=X2|v1_subset_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_36, plain, ![X28, X29]:((~m1_subset_1(X28,k1_zfmisc_1(X29))|r1_tarski(X28,X29))&(~r1_tarski(X28,X29)|m1_subset_1(X28,k1_zfmisc_1(X29)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_37, plain, (r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_38, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),X1)|r2_hidden(esk5_2(u1_struct_0(esk3_0),X1),esk4_0)|~r1_orders_2(esk3_0,X2,esk5_2(u1_struct_0(esk3_0),X1))|~r2_hidden(X2,esk4_0)), inference(spm,[status(thm)],[c_0_30, c_0_28])).
cnf(c_0_39, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),X1)|r1_orders_2(esk3_0,k3_yellow_0(esk3_0),esk5_2(u1_struct_0(esk3_0),X1))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33]), c_0_26])]), c_0_34])).
cnf(c_0_40, negated_conjecture, (r2_hidden(k3_yellow_0(esk3_0),esk4_0)|~v1_subset_1(esk4_0,u1_struct_0(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_41, negated_conjecture, (u1_struct_0(esk3_0)=esk4_0|v1_subset_1(esk4_0,u1_struct_0(esk3_0))), inference(spm,[status(thm)],[c_0_35, c_0_24])).
fof(c_0_42, plain, ![X30, X31, X32]:((m1_subset_1(esk6_3(X30,X31,X32),X30)|X31=X32|~m1_subset_1(X32,k1_zfmisc_1(X30))|~m1_subset_1(X31,k1_zfmisc_1(X30)))&((~r2_hidden(esk6_3(X30,X31,X32),X31)|~r2_hidden(esk6_3(X30,X31,X32),X32)|X31=X32|~m1_subset_1(X32,k1_zfmisc_1(X30))|~m1_subset_1(X31,k1_zfmisc_1(X30)))&(r2_hidden(esk6_3(X30,X31,X32),X31)|r2_hidden(esk6_3(X30,X31,X32),X32)|X31=X32|~m1_subset_1(X32,k1_zfmisc_1(X30))|~m1_subset_1(X31,k1_zfmisc_1(X30))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_subset_1])])])])])).
cnf(c_0_43, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_37, c_0_22])).
fof(c_0_45, plain, ![X34, X35, X36]:(~m1_subset_1(X35,k1_zfmisc_1(X34))|(~r2_hidden(X36,X35)|r2_hidden(X36,X34))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
cnf(c_0_46, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),X1)|r2_hidden(esk5_2(u1_struct_0(esk3_0),X1),esk4_0)|~r2_hidden(k3_yellow_0(esk3_0),esk4_0)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_47, negated_conjecture, (u1_struct_0(esk3_0)=esk4_0|r2_hidden(k3_yellow_0(esk3_0),esk4_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_48, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r2_hidden(esk6_3(X1,X2,X3),X3)|X2=X3|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_49, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_50, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_51, negated_conjecture, (u1_struct_0(esk3_0)=esk4_0|r1_tarski(u1_struct_0(esk3_0),X1)|r2_hidden(esk5_2(u1_struct_0(esk3_0),X1),esk4_0)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_52, plain, (X1=X2|r2_hidden(esk6_3(X2,X1,X2),X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])).
fof(c_0_53, plain, ![X37]:(~l1_orders_2(X37)|m1_subset_1(k3_yellow_0(X37),u1_struct_0(X37))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_yellow_0])])).
cnf(c_0_54, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_55, negated_conjecture, (u1_struct_0(esk3_0)=esk4_0|r1_tarski(u1_struct_0(esk3_0),esk4_0)), inference(spm,[status(thm)],[c_0_37, c_0_51])).
cnf(c_0_56, plain, (X2=X3|~r2_hidden(esk6_3(X1,X2,X3),X2)|~r2_hidden(esk6_3(X1,X2,X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_57, negated_conjecture, (u1_struct_0(esk3_0)=esk4_0|r2_hidden(esk6_3(u1_struct_0(esk3_0),esk4_0,u1_struct_0(esk3_0)),u1_struct_0(esk3_0))), inference(spm,[status(thm)],[c_0_52, c_0_24])).
cnf(c_0_58, plain, (~v1_subset_1(X1,X2)|X1!=X2|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_59, plain, ![X38, X39]:(~m1_subset_1(X38,X39)|(v1_xboole_0(X39)|r2_hidden(X38,X39))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_60, plain, (m1_subset_1(k3_yellow_0(X1),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_61, negated_conjecture, (u1_struct_0(esk3_0)=esk4_0|r2_hidden(X1,esk4_0)|~r2_hidden(X1,u1_struct_0(esk3_0))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_62, negated_conjecture, (u1_struct_0(esk3_0)=esk4_0|~r2_hidden(esk6_3(u1_struct_0(esk3_0),esk4_0,u1_struct_0(esk3_0)),esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_49]), c_0_24])])).
cnf(c_0_63, plain, (~v1_subset_1(X1,X1)|~m1_subset_1(X1,k1_zfmisc_1(X1))), inference(er,[status(thm)],[c_0_58])).
cnf(c_0_64, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_65, negated_conjecture, (m1_subset_1(k3_yellow_0(esk3_0),u1_struct_0(esk3_0))), inference(spm,[status(thm)],[c_0_60, c_0_26])).
cnf(c_0_66, negated_conjecture, (v1_subset_1(esk4_0,u1_struct_0(esk3_0))|~r2_hidden(k3_yellow_0(esk3_0),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_67, negated_conjecture, (u1_struct_0(esk3_0)=esk4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_57]), c_0_62])).
cnf(c_0_68, plain, (~v1_subset_1(X1,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_49])])).
cnf(c_0_69, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r2_hidden(k3_yellow_0(esk3_0),u1_struct_0(esk3_0))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_70, negated_conjecture, (~v1_xboole_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_71, negated_conjecture, (~r2_hidden(k3_yellow_0(esk3_0),esk4_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_67]), c_0_68])).
cnf(c_0_72, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_67]), c_0_67]), c_0_70]), c_0_71]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 73
# Proof object clause steps            : 48
# Proof object formula steps           : 25
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 34
# Processed clauses                    : 202
# ...of these trivial                  : 0
# ...subsumed                          : 39
# ...remaining for further processing  : 162
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 12
# Backward-rewritten                   : 57
# Generated clauses                    : 191
# ...of the previous two non-trivial   : 189
# Contextual simplify-reflections      : 5
# Paramodulations                      : 190
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 58
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 43
# Current number of unprocessed clauses: 53
# ...number of literals in the above   : 170
# Current number of archived formulas  : 0
# Current number of archived clauses   : 103
# Clause-clause subsumption calls (NU) : 1301
# Rec. Clause-clause subsumption calls : 667
# Non-unit clause-clause subsumptions  : 56
# Unit Clause-clause subsumption calls : 44
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 13
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6449

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.000 s
# Total time               : 0.015 s
# Maximum resident set size: 28824 pages
