# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t64_orders_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_orders_2(X1,X2,X3)=>r1_tarski(k3_orders_2(X1,X4,X2),k3_orders_2(X1,X4,X3))))))), file('data/MPTP/proofs/5168d0b2a77643f3bb358ab2d92ebe55.E_input', t64_orders_2)).
fof(t57_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X2,k3_orders_2(X1,X4,X3))<=>(r2_orders_2(X1,X2,X3)&r2_hidden(X2,X4))))))), file('data/MPTP/proofs/5168d0b2a77643f3bb358ab2d92ebe55.E_input', t57_orders_2)).
fof(t7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,X2)=>r2_hidden(X4,X3)))=>r1_tarski(X2,X3)))), file('data/MPTP/proofs/5168d0b2a77643f3bb358ab2d92ebe55.E_input', t7_subset_1)).
fof(t32_orders_2, axiom, ![X1]:((v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((r2_orders_2(X1,X2,X3)&r1_orders_2(X1,X3,X4))|(r1_orders_2(X1,X2,X3)&r2_orders_2(X1,X3,X4)))=>r2_orders_2(X1,X2,X4)))))), file('data/MPTP/proofs/5168d0b2a77643f3bb358ab2d92ebe55.E_input', t32_orders_2)).
fof(d10_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_orders_2(X1,X2,X3)<=>(r1_orders_2(X1,X2,X3)&~(X2=X3)))))), file('data/MPTP/proofs/5168d0b2a77643f3bb358ab2d92ebe55.E_input', d10_orders_2)).
fof(dt_k3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))&(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k3_orders_2(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/5168d0b2a77643f3bb358ab2d92ebe55.E_input', dt_k3_orders_2)).
fof(c_0_6, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_orders_2(X1,X2,X3)=>r1_tarski(k3_orders_2(X1,X4,X2),k3_orders_2(X1,X4,X3)))))))), inference(assume_negation,[status(cth)],[t64_orders_2])).
fof(c_0_7, plain, ![X9, X10, X11, X12]:(((r2_orders_2(X9,X10,X11)|~r2_hidden(X10,k3_orders_2(X9,X12,X11))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X9)))|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(v2_struct_0(X9)|(~v3_orders_2(X9)|(~v4_orders_2(X9)|(~v5_orders_2(X9)|~l1_orders_2(X9))))))&(r2_hidden(X10,X12)|~r2_hidden(X10,k3_orders_2(X9,X12,X11))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X9)))|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(v2_struct_0(X9)|(~v3_orders_2(X9)|(~v4_orders_2(X9)|(~v5_orders_2(X9)|~l1_orders_2(X9)))))))&(~r2_orders_2(X9,X10,X11)|~r2_hidden(X10,X12)|r2_hidden(X10,k3_orders_2(X9,X12,X11))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X9)))|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(v2_struct_0(X9)|(~v3_orders_2(X9)|(~v4_orders_2(X9)|(~v5_orders_2(X9)|~l1_orders_2(X9))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t57_orders_2])])])])])).
fof(c_0_8, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(r2_orders_2(esk1_0,esk2_0,esk3_0)&~r1_tarski(k3_orders_2(esk1_0,esk4_0,esk2_0),k3_orders_2(esk1_0,esk4_0,esk3_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_6])])])])).
fof(c_0_9, plain, ![X23, X24, X25]:((m1_subset_1(esk5_3(X23,X24,X25),X23)|r1_tarski(X24,X25)|~m1_subset_1(X25,k1_zfmisc_1(X23))|~m1_subset_1(X24,k1_zfmisc_1(X23)))&((r2_hidden(esk5_3(X23,X24,X25),X24)|r1_tarski(X24,X25)|~m1_subset_1(X25,k1_zfmisc_1(X23))|~m1_subset_1(X24,k1_zfmisc_1(X23)))&(~r2_hidden(esk5_3(X23,X24,X25),X25)|r1_tarski(X24,X25)|~m1_subset_1(X25,k1_zfmisc_1(X23))|~m1_subset_1(X24,k1_zfmisc_1(X23))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_subset_1])])])])])).
cnf(c_0_10, plain, (r2_hidden(X2,k3_orders_2(X1,X4,X3))|v2_struct_0(X1)|~r2_orders_2(X1,X2,X3)|~r2_hidden(X2,X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_17, plain, ![X16, X17, X18, X19]:((~r2_orders_2(X16,X17,X18)|~r1_orders_2(X16,X18,X19)|r2_orders_2(X16,X17,X19)|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(~v4_orders_2(X16)|(~v5_orders_2(X16)|~l1_orders_2(X16))))&(~r1_orders_2(X16,X17,X18)|~r2_orders_2(X16,X18,X19)|r2_orders_2(X16,X17,X19)|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(~v4_orders_2(X16)|(~v5_orders_2(X16)|~l1_orders_2(X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t32_orders_2])])])])).
fof(c_0_18, plain, ![X13, X14, X15]:(((r1_orders_2(X13,X14,X15)|~r2_orders_2(X13,X14,X15)|~m1_subset_1(X15,u1_struct_0(X13))|~m1_subset_1(X14,u1_struct_0(X13))|~l1_orders_2(X13))&(X14!=X15|~r2_orders_2(X13,X14,X15)|~m1_subset_1(X15,u1_struct_0(X13))|~m1_subset_1(X14,u1_struct_0(X13))|~l1_orders_2(X13)))&(~r1_orders_2(X13,X14,X15)|X14=X15|r2_orders_2(X13,X14,X15)|~m1_subset_1(X15,u1_struct_0(X13))|~m1_subset_1(X14,u1_struct_0(X13))|~l1_orders_2(X13))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_orders_2])])])])])).
cnf(c_0_19, plain, (r1_tarski(X2,X3)|~r2_hidden(esk5_3(X1,X2,X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_20, negated_conjecture, (r2_hidden(X1,k3_orders_2(esk1_0,esk4_0,X2))|~r2_hidden(X1,esk4_0)|~r2_orders_2(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12]), c_0_13]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_21, plain, (r2_hidden(X1,X2)|v2_struct_0(X3)|~r2_hidden(X1,k3_orders_2(X3,X2,X4))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X4,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~v3_orders_2(X3)|~v4_orders_2(X3)|~v5_orders_2(X3)|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_22, plain, (r2_hidden(esk5_3(X1,X2,X3),X2)|r1_tarski(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_23, plain, (r2_orders_2(X1,X2,X4)|~r1_orders_2(X1,X2,X3)|~r2_orders_2(X1,X3,X4)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (r1_orders_2(X1,X2,X3)|~r2_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, negated_conjecture, (r1_tarski(X1,k3_orders_2(esk1_0,esk4_0,X2))|~r2_hidden(esk5_3(X3,X1,k3_orders_2(esk1_0,esk4_0,X2)),esk4_0)|~r2_orders_2(esk1_0,esk5_3(X3,X1,k3_orders_2(esk1_0,esk4_0,X2)),X2)|~m1_subset_1(esk5_3(X3,X1,k3_orders_2(esk1_0,esk4_0,X2)),u1_struct_0(esk1_0))|~m1_subset_1(k3_orders_2(esk1_0,esk4_0,X2),k1_zfmisc_1(X3))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_26, plain, (r2_hidden(esk5_3(X1,k3_orders_2(X2,X3,X4),X5),X3)|r1_tarski(k3_orders_2(X2,X3,X4),X5)|v2_struct_0(X2)|~m1_subset_1(esk5_3(X1,k3_orders_2(X2,X3,X4),X5),u1_struct_0(X2))|~m1_subset_1(k3_orders_2(X2,X3,X4),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(X1))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27, plain, (r2_orders_2(X1,X2,X3)|~r2_orders_2(X1,X4,X3)|~r2_orders_2(X1,X2,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_28, negated_conjecture, (r2_orders_2(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_30, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_31, negated_conjecture, (r1_tarski(k3_orders_2(X1,esk4_0,X2),k3_orders_2(esk1_0,esk4_0,X3))|v2_struct_0(X1)|~r2_orders_2(esk1_0,esk5_3(X4,k3_orders_2(X1,esk4_0,X2),k3_orders_2(esk1_0,esk4_0,X3)),X3)|~m1_subset_1(esk5_3(X4,k3_orders_2(X1,esk4_0,X2),k3_orders_2(esk1_0,esk4_0,X3)),u1_struct_0(esk1_0))|~m1_subset_1(esk5_3(X4,k3_orders_2(X1,esk4_0,X2),k3_orders_2(esk1_0,esk4_0,X3)),u1_struct_0(X1))|~m1_subset_1(k3_orders_2(esk1_0,esk4_0,X3),k1_zfmisc_1(X4))|~m1_subset_1(k3_orders_2(X1,esk4_0,X2),k1_zfmisc_1(X4))|~m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_32, negated_conjecture, (r2_orders_2(esk1_0,X1,esk3_0)|~r2_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30]), c_0_12]), c_0_13]), c_0_14])])).
cnf(c_0_33, plain, (r2_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r2_hidden(X2,k3_orders_2(X1,X4,X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_34, negated_conjecture, (r1_tarski(k3_orders_2(X1,esk4_0,X2),k3_orders_2(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r2_orders_2(esk1_0,esk5_3(X3,k3_orders_2(X1,esk4_0,X2),k3_orders_2(esk1_0,esk4_0,esk3_0)),esk2_0)|~m1_subset_1(esk5_3(X3,k3_orders_2(X1,esk4_0,X2),k3_orders_2(esk1_0,esk4_0,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk5_3(X3,k3_orders_2(X1,esk4_0,X2),k3_orders_2(esk1_0,esk4_0,esk3_0)),u1_struct_0(X1))|~m1_subset_1(k3_orders_2(esk1_0,esk4_0,esk3_0),k1_zfmisc_1(X3))|~m1_subset_1(k3_orders_2(X1,esk4_0,X2),k1_zfmisc_1(X3))|~m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_29])])).
cnf(c_0_35, plain, (r1_tarski(k3_orders_2(X1,X2,X3),X4)|r2_orders_2(X1,esk5_3(X5,k3_orders_2(X1,X2,X3),X4),X3)|v2_struct_0(X1)|~m1_subset_1(esk5_3(X5,k3_orders_2(X1,X2,X3),X4),u1_struct_0(X1))|~m1_subset_1(k3_orders_2(X1,X2,X3),k1_zfmisc_1(X5))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(X5))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_33, c_0_22])).
cnf(c_0_36, negated_conjecture, (~r1_tarski(k3_orders_2(esk1_0,esk4_0,esk2_0),k3_orders_2(esk1_0,esk4_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_37, plain, ![X20, X21, X22]:(v2_struct_0(X20)|(~v3_orders_2(X20)|(~v4_orders_2(X20)|(~v5_orders_2(X20)|~l1_orders_2(X20))))|(~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X22,u1_struct_0(X20)))|m1_subset_1(k3_orders_2(X20,X21,X22),k1_zfmisc_1(u1_struct_0(X20)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_orders_2])])])).
cnf(c_0_38, negated_conjecture, (~m1_subset_1(esk5_3(X1,k3_orders_2(esk1_0,esk4_0,esk2_0),k3_orders_2(esk1_0,esk4_0,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(k3_orders_2(esk1_0,esk4_0,esk3_0),k1_zfmisc_1(X1))|~m1_subset_1(k3_orders_2(esk1_0,esk4_0,esk2_0),k1_zfmisc_1(X1))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_11]), c_0_30]), c_0_12]), c_0_13]), c_0_14]), c_0_15])]), c_0_36]), c_0_16])).
cnf(c_0_39, plain, (m1_subset_1(esk5_3(X1,X2,X3),X1)|r1_tarski(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_40, plain, (v2_struct_0(X1)|m1_subset_1(k3_orders_2(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_41, negated_conjecture, (~m1_subset_1(k3_orders_2(esk1_0,esk4_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k3_orders_2(esk1_0,esk4_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_36])).
cnf(c_0_42, negated_conjecture, (m1_subset_1(k3_orders_2(esk1_0,esk4_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_11]), c_0_12]), c_0_13]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_43, negated_conjecture, (~m1_subset_1(k3_orders_2(esk1_0,esk4_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_29])])).
cnf(c_0_44, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_42]), c_0_30])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 45
# Proof object clause steps            : 32
# Proof object formula steps           : 13
# Proof object conjectures             : 23
# Proof object clause conjectures      : 20
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 34
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 22
# Processed clauses                    : 101
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 99
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 97
# ...of the previous two non-trivial   : 81
# Contextual simplify-reflections      : 2
# Paramodulations                      : 96
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 75
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 62
# Current number of unprocessed clauses: 24
# ...number of literals in the above   : 300
# Current number of archived formulas  : 0
# Current number of archived clauses   : 23
# Clause-clause subsumption calls (NU) : 2956
# Rec. Clause-clause subsumption calls : 113
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 43
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 9511

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.000 s
# Total time               : 0.015 s
# Maximum resident set size: 28788 pages
