# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.011 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', d10_xboole_0)).
fof(t4_tsep_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X1)=>(r1_tarski(u1_struct_0(X2),u1_struct_0(X3))<=>m1_pre_topc(X2,X3))))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', t4_tsep_1)).
fof(t116_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>((v1_tsep_1(X2,X1)&m1_pre_topc(X2,X1))<=>g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=k8_tmap_1(X1,X2)))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', t116_tmap_1)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', d4_subset_1)).
fof(t39_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', t39_pre_topc)).
fof(t1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', t1_tsep_1)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', dt_m1_pre_topc)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', t4_subset)).
fof(d5_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>r2_hidden(X2,u1_pre_topc(X1))))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', d5_pre_topc)).
fof(t103_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X2,u1_pre_topc(X1))<=>u1_pre_topc(X1)=k5_tmap_1(X1,X2)))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', t103_tmap_1)).
fof(cc1_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>v2_pre_topc(X2))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', cc1_pre_topc)).
fof(d11_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:((v1_pre_topc(X3)&(v2_pre_topc(X3)&l1_pre_topc(X3)))=>(X3=k8_tmap_1(X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(X4=u1_struct_0(X2)=>X3=k6_tmap_1(X1,X4))))))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', d11_tmap_1)).
fof(dt_k8_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_pre_topc(X2,X1))=>(v1_pre_topc(k8_tmap_1(X1,X2))&(v2_pre_topc(k8_tmap_1(X1,X2))&l1_pre_topc(k8_tmap_1(X1,X2))))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', dt_k8_tmap_1)).
fof(t104_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(u1_struct_0(k6_tmap_1(X1,X2))=u1_struct_0(X1)&u1_pre_topc(k6_tmap_1(X1,X2))=k5_tmap_1(X1,X2)))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', t104_tmap_1)).
fof(t11_tmap_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>(v1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))&m1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)),X1)))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', t11_tmap_1)).
fof(d1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>(v1_tsep_1(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>v3_pre_topc(X3,X1)))))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', d1_tsep_1)).
fof(t12_tmap_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:((v2_pre_topc(X2)&l1_pre_topc(X2))=>![X3]:((v2_pre_topc(X3)&l1_pre_topc(X3))=>(X2=g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))=>(m1_pre_topc(X2,X1)<=>m1_pre_topc(X3,X1)))))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', t12_tmap_1)).
fof(t58_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))=>v2_pre_topc(X1))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', t58_pre_topc)).
fof(t2_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>m1_pre_topc(X1,X1)), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', t2_tsep_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', d2_subset_1)).
fof(t65_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>(m1_pre_topc(X1,X2)<=>m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', t65_pre_topc)).
fof(t7_tsep_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X2)=>m1_pre_topc(X3,X1)))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', t7_tsep_1)).
fof(abstractness_v1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v1_pre_topc(X1)=>X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', abstractness_v1_pre_topc)).
fof(t114_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>(u1_struct_0(k8_tmap_1(X1,X2))=u1_struct_0(X1)&![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>u1_pre_topc(k8_tmap_1(X1,X2))=k5_tmap_1(X1,X3)))))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', t114_tmap_1)).
fof(t19_tsep_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((v1_tsep_1(X2,X1)&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(r1_tarski(u1_struct_0(X2),u1_struct_0(X3))=>(v1_tsep_1(X2,X3)&m1_pre_topc(X2,X3)))))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', t19_tsep_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', t3_subset)).
fof(dt_k6_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_pre_topc(k6_tmap_1(X1,X2))&(v2_pre_topc(k6_tmap_1(X1,X2))&l1_pre_topc(k6_tmap_1(X1,X2))))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', dt_k6_tmap_1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', fc2_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', dt_l1_pre_topc)).
fof(fc1_struct_0, axiom, ![X1]:((v2_struct_0(X1)&l1_struct_0(X1))=>v1_xboole_0(u1_struct_0(X1))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', fc1_struct_0)).
fof(fc6_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>(v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))&v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', fc6_pre_topc)).
fof(t105_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k6_tmap_1(X1,X2))))=>(X3=X2=>v3_pre_topc(X3,k6_tmap_1(X1,X2)))))), file('data/MPTP/proofs/516e91b8a9ca4021b96cf2399b21f2b5.E_input', t105_tmap_1)).
fof(c_0_33, plain, ![X71, X72]:(((r1_tarski(X71,X72)|X71!=X72)&(r1_tarski(X72,X71)|X71!=X72))&(~r1_tarski(X71,X72)|~r1_tarski(X72,X71)|X71=X72)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_34, plain, ![X22, X23, X24]:((~r1_tarski(u1_struct_0(X23),u1_struct_0(X24))|m1_pre_topc(X23,X24)|~m1_pre_topc(X24,X22)|~m1_pre_topc(X23,X22)|(~v2_pre_topc(X22)|~l1_pre_topc(X22)))&(~m1_pre_topc(X23,X24)|r1_tarski(u1_struct_0(X23),u1_struct_0(X24))|~m1_pre_topc(X24,X22)|~m1_pre_topc(X23,X22)|(~v2_pre_topc(X22)|~l1_pre_topc(X22)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_tsep_1])])])])).
cnf(c_0_35, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_36, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>((v1_tsep_1(X2,X1)&m1_pre_topc(X2,X1))<=>g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=k8_tmap_1(X1,X2))))), inference(assume_negation,[status(cth)],[t116_tmap_1])).
fof(c_0_37, plain, ![X37]:m1_subset_1(k2_subset_1(X37),k1_zfmisc_1(X37)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_38, plain, ![X25]:k2_subset_1(X25)=X25, inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_39, plain, ![X28, X29, X30]:(~l1_pre_topc(X28)|(~m1_pre_topc(X29,X28)|(~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X29)))|m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X28)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t39_pre_topc])])])).
fof(c_0_40, plain, ![X73, X74]:(~l1_pre_topc(X73)|(~m1_pre_topc(X74,X73)|m1_subset_1(u1_struct_0(X74),k1_zfmisc_1(u1_struct_0(X73))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])).
fof(c_0_41, plain, ![X51, X52]:(~l1_pre_topc(X51)|(~m1_pre_topc(X52,X51)|l1_pre_topc(X52))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_42, plain, (m1_pre_topc(X1,X2)|~r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_35])).
fof(c_0_44, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((~v1_tsep_1(esk2_0,esk1_0)|~m1_pre_topc(esk2_0,esk1_0)|g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))!=k8_tmap_1(esk1_0,esk2_0))&((v1_tsep_1(esk2_0,esk1_0)|g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))=k8_tmap_1(esk1_0,esk2_0))&(m1_pre_topc(esk2_0,esk1_0)|g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))=k8_tmap_1(esk1_0,esk2_0)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_36])])])])])).
fof(c_0_45, plain, ![X8, X9, X10]:(~r2_hidden(X8,X9)|~m1_subset_1(X9,k1_zfmisc_1(X10))|m1_subset_1(X8,X10)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_46, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_47, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_48, plain, ![X16, X17]:((~v3_pre_topc(X17,X16)|r2_hidden(X17,u1_pre_topc(X16))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~l1_pre_topc(X16))&(~r2_hidden(X17,u1_pre_topc(X16))|v3_pre_topc(X17,X16)|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~l1_pre_topc(X16))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_pre_topc])])])])).
fof(c_0_49, plain, ![X61, X62]:((~r2_hidden(X62,u1_pre_topc(X61))|u1_pre_topc(X61)=k5_tmap_1(X61,X62)|~m1_subset_1(X62,k1_zfmisc_1(u1_struct_0(X61)))|(v2_struct_0(X61)|(~v2_pre_topc(X61)|~l1_pre_topc(X61))))&(u1_pre_topc(X61)!=k5_tmap_1(X61,X62)|r2_hidden(X62,u1_pre_topc(X61))|~m1_subset_1(X62,k1_zfmisc_1(u1_struct_0(X61)))|(v2_struct_0(X61)|(~v2_pre_topc(X61)|~l1_pre_topc(X61))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t103_tmap_1])])])])])).
cnf(c_0_50, plain, (m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_52, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_53, plain, (m1_pre_topc(X1,X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_54, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_55, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_56, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_57, plain, ![X43, X44]:(~v2_pre_topc(X43)|~l1_pre_topc(X43)|(~m1_pre_topc(X44,X43)|v2_pre_topc(X44))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_pre_topc])])])).
fof(c_0_58, plain, ![X82, X83, X84, X85]:((X84!=k8_tmap_1(X82,X83)|(~m1_subset_1(X85,k1_zfmisc_1(u1_struct_0(X82)))|(X85!=u1_struct_0(X83)|X84=k6_tmap_1(X82,X85)))|(~v1_pre_topc(X84)|(~v2_pre_topc(X84)|~l1_pre_topc(X84)))|~m1_pre_topc(X83,X82)|(v2_struct_0(X82)|(~v2_pre_topc(X82)|~l1_pre_topc(X82))))&((m1_subset_1(esk5_3(X82,X83,X84),k1_zfmisc_1(u1_struct_0(X82)))|X84=k8_tmap_1(X82,X83)|(~v1_pre_topc(X84)|(~v2_pre_topc(X84)|~l1_pre_topc(X84)))|~m1_pre_topc(X83,X82)|(v2_struct_0(X82)|(~v2_pre_topc(X82)|~l1_pre_topc(X82))))&((esk5_3(X82,X83,X84)=u1_struct_0(X83)|X84=k8_tmap_1(X82,X83)|(~v1_pre_topc(X84)|(~v2_pre_topc(X84)|~l1_pre_topc(X84)))|~m1_pre_topc(X83,X82)|(v2_struct_0(X82)|(~v2_pre_topc(X82)|~l1_pre_topc(X82))))&(X84!=k6_tmap_1(X82,esk5_3(X82,X83,X84))|X84=k8_tmap_1(X82,X83)|(~v1_pre_topc(X84)|(~v2_pre_topc(X84)|~l1_pre_topc(X84)))|~m1_pre_topc(X83,X82)|(v2_struct_0(X82)|(~v2_pre_topc(X82)|~l1_pre_topc(X82))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d11_tmap_1])])])])])])).
fof(c_0_59, plain, ![X18, X19]:((v1_pre_topc(k8_tmap_1(X18,X19))|(v2_struct_0(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18))|~m1_pre_topc(X19,X18)))&((v2_pre_topc(k8_tmap_1(X18,X19))|(v2_struct_0(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18))|~m1_pre_topc(X19,X18)))&(l1_pre_topc(k8_tmap_1(X18,X19))|(v2_struct_0(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18))|~m1_pre_topc(X19,X18))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k8_tmap_1])])])])).
cnf(c_0_60, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_61, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_62, plain, (v3_pre_topc(X1,X2)|~r2_hidden(X1,u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_63, plain, (r2_hidden(X2,u1_pre_topc(X1))|v2_struct_0(X1)|u1_pre_topc(X1)!=k5_tmap_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_64, plain, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52])).
cnf(c_0_65, negated_conjecture, (m1_pre_topc(esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55]), c_0_56])])).
cnf(c_0_66, negated_conjecture, (l1_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_54]), c_0_55])])).
cnf(c_0_67, plain, (v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
fof(c_0_68, plain, ![X20, X21]:((u1_struct_0(k6_tmap_1(X20,X21))=u1_struct_0(X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20))))&(u1_pre_topc(k6_tmap_1(X20,X21))=k5_tmap_1(X20,X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t104_tmap_1])])])])])).
cnf(c_0_69, plain, (X1=k6_tmap_1(X2,X4)|v2_struct_0(X2)|X1!=k8_tmap_1(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|X4!=u1_struct_0(X3)|~v1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_70, plain, (v2_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_71, plain, (l1_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_72, plain, (v1_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_73, plain, (u1_pre_topc(X2)=k5_tmap_1(X2,X1)|v2_struct_0(X2)|~r2_hidden(X1,u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_74, plain, (r2_hidden(X1,u1_pre_topc(X2))|~v3_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_75, plain, ![X26, X27]:((v1_pre_topc(g1_pre_topc(u1_struct_0(X27),u1_pre_topc(X27)))|~m1_pre_topc(X27,X26)|~l1_pre_topc(X26))&(m1_pre_topc(g1_pre_topc(u1_struct_0(X27),u1_pre_topc(X27)),X26)|~m1_pre_topc(X27,X26)|~l1_pre_topc(X26))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_tmap_1])])])])).
cnf(c_0_76, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_77, plain, (v3_pre_topc(X1,X2)|v2_struct_0(X2)|k5_tmap_1(X2,X1)!=u1_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_78, negated_conjecture, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_66])])).
cnf(c_0_79, negated_conjecture, (v2_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_54]), c_0_55]), c_0_56])])).
cnf(c_0_80, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_81, plain, (u1_pre_topc(k6_tmap_1(X1,X2))=k5_tmap_1(X1,X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_82, plain, (k6_tmap_1(X1,u1_struct_0(X2))=k8_tmap_1(X1,X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_69])]), c_0_70]), c_0_71]), c_0_72]), c_0_51])).
cnf(c_0_83, plain, (k5_tmap_1(X1,X2)=u1_pre_topc(X1)|v2_struct_0(X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
fof(c_0_84, plain, ![X38, X39, X40]:((~v1_tsep_1(X39,X38)|(~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X38)))|(X40!=u1_struct_0(X39)|v3_pre_topc(X40,X38)))|~m1_pre_topc(X39,X38)|~l1_pre_topc(X38))&((m1_subset_1(esk3_2(X38,X39),k1_zfmisc_1(u1_struct_0(X38)))|v1_tsep_1(X39,X38)|~m1_pre_topc(X39,X38)|~l1_pre_topc(X38))&((esk3_2(X38,X39)=u1_struct_0(X39)|v1_tsep_1(X39,X38)|~m1_pre_topc(X39,X38)|~l1_pre_topc(X38))&(~v3_pre_topc(esk3_2(X38,X39),X38)|v1_tsep_1(X39,X38)|~m1_pre_topc(X39,X38)|~l1_pre_topc(X38))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tsep_1])])])])])).
fof(c_0_85, plain, ![X13, X14, X15]:((~m1_pre_topc(X14,X13)|m1_pre_topc(X15,X13)|X14!=g1_pre_topc(u1_struct_0(X15),u1_pre_topc(X15))|(~v2_pre_topc(X15)|~l1_pre_topc(X15))|(~v2_pre_topc(X14)|~l1_pre_topc(X14))|~l1_pre_topc(X13))&(~m1_pre_topc(X15,X13)|m1_pre_topc(X14,X13)|X14!=g1_pre_topc(u1_struct_0(X15),u1_pre_topc(X15))|(~v2_pre_topc(X15)|~l1_pre_topc(X15))|(~v2_pre_topc(X14)|~l1_pre_topc(X14))|~l1_pre_topc(X13))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_tmap_1])])])])).
fof(c_0_86, plain, ![X42]:(~l1_pre_topc(X42)|(~v2_pre_topc(g1_pre_topc(u1_struct_0(X42),u1_pre_topc(X42)))|v2_pre_topc(X42))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t58_pre_topc])])).
cnf(c_0_87, plain, (m1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_75])).
fof(c_0_88, plain, ![X60]:(~l1_pre_topc(X60)|m1_pre_topc(X60,X60)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tsep_1])])).
cnf(c_0_89, plain, (m1_subset_1(X1,u1_pre_topc(X2))|~v3_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_76, c_0_74])).
cnf(c_0_90, negated_conjecture, (v3_pre_topc(u1_struct_0(X1),esk2_0)|k5_tmap_1(esk2_0,u1_struct_0(X1))!=u1_pre_topc(esk2_0)|~m1_pre_topc(X1,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_66]), c_0_79])]), c_0_80])).
cnf(c_0_91, plain, (k5_tmap_1(X1,u1_struct_0(X2))=u1_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_51])).
cnf(c_0_92, negated_conjecture, (k5_tmap_1(esk2_0,u1_struct_0(X1))=u1_pre_topc(esk2_0)|~v3_pre_topc(u1_struct_0(X1),esk2_0)|~m1_pre_topc(X1,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_78]), c_0_66]), c_0_79])]), c_0_80])).
cnf(c_0_93, plain, (v3_pre_topc(X3,X2)|~v1_tsep_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|X3!=u1_struct_0(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_84])).
fof(c_0_94, plain, ![X55, X56]:(((~m1_subset_1(X56,X55)|r2_hidden(X56,X55)|v1_xboole_0(X55))&(~r2_hidden(X56,X55)|m1_subset_1(X56,X55)|v1_xboole_0(X55)))&((~m1_subset_1(X56,X55)|v1_xboole_0(X56)|~v1_xboole_0(X55))&(~v1_xboole_0(X56)|m1_subset_1(X56,X55)|~v1_xboole_0(X55)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
cnf(c_0_95, plain, (m1_subset_1(X1,u1_pre_topc(X2))|v2_struct_0(X2)|k5_tmap_1(X2,X1)!=u1_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_76, c_0_63])).
cnf(c_0_96, plain, (m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)|X1!=g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_97, plain, (v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_98, plain, (l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_52, c_0_87])).
cnf(c_0_99, plain, (m1_pre_topc(X1,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_100, plain, (m1_subset_1(u1_struct_0(X1),u1_pre_topc(X1))|~v3_pre_topc(u1_struct_0(X1),X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_89, c_0_61])).
cnf(c_0_101, negated_conjecture, (v3_pre_topc(u1_struct_0(X1),esk2_0)|u1_pre_topc(k8_tmap_1(esk2_0,X1))!=u1_pre_topc(esk2_0)|~m1_pre_topc(X1,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_66]), c_0_79])]), c_0_80])).
cnf(c_0_102, negated_conjecture, (u1_pre_topc(k8_tmap_1(esk2_0,X1))=u1_pre_topc(esk2_0)|~v3_pre_topc(u1_struct_0(X1),esk2_0)|~m1_pre_topc(X1,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_91]), c_0_66]), c_0_79])]), c_0_80])).
cnf(c_0_103, plain, (v3_pre_topc(u1_struct_0(X1),X2)|~v1_tsep_1(X1,X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_93]), c_0_51])).
cnf(c_0_104, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_105, negated_conjecture, (m1_subset_1(u1_struct_0(X1),u1_pre_topc(esk2_0))|k5_tmap_1(esk2_0,u1_struct_0(X1))!=u1_pre_topc(esk2_0)|~m1_pre_topc(X1,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_78]), c_0_66]), c_0_79])]), c_0_80])).
fof(c_0_106, plain, ![X53, X54]:((~m1_pre_topc(X53,X54)|m1_pre_topc(X53,g1_pre_topc(u1_struct_0(X54),u1_pre_topc(X54)))|~l1_pre_topc(X54)|~l1_pre_topc(X53))&(~m1_pre_topc(X53,g1_pre_topc(u1_struct_0(X54),u1_pre_topc(X54)))|m1_pre_topc(X53,X54)|~l1_pre_topc(X54)|~l1_pre_topc(X53))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t65_pre_topc])])])])).
cnf(c_0_107, plain, (m1_pre_topc(X1,X2)|~m1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),X2)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_96, c_0_52])]), c_0_97])).
cnf(c_0_108, plain, (l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_109, plain, (v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_67, c_0_87])).
cnf(c_0_110, negated_conjecture, (m1_subset_1(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))|u1_pre_topc(k8_tmap_1(esk2_0,esk2_0))!=u1_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_66]), c_0_65])])).
cnf(c_0_111, negated_conjecture, (u1_pre_topc(k8_tmap_1(esk2_0,X1))=u1_pre_topc(esk2_0)|~v1_tsep_1(X1,esk2_0)|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_66])])).
cnf(c_0_112, plain, (v3_pre_topc(X1,X2)|v1_xboole_0(u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_pre_topc(X2))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_62, c_0_104])).
cnf(c_0_113, negated_conjecture, (m1_subset_1(u1_struct_0(X1),u1_pre_topc(esk2_0))|~v3_pre_topc(u1_struct_0(X1),esk2_0)|~m1_pre_topc(X1,esk2_0)), inference(spm,[status(thm)],[c_0_105, c_0_92])).
cnf(c_0_114, plain, (m1_pre_topc(X1,X2)|~m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_115, plain, (m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_99]), c_0_108])).
cnf(c_0_116, negated_conjecture, (v2_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_65]), c_0_66]), c_0_79])])).
cnf(c_0_117, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_118, negated_conjecture, (m1_subset_1(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))|~v1_tsep_1(esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_65])])).
cnf(c_0_119, plain, (v3_pre_topc(u1_struct_0(X1),X1)|v1_xboole_0(u1_pre_topc(X1))|~m1_subset_1(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_112, c_0_61])).
cnf(c_0_120, negated_conjecture, (m1_subset_1(u1_struct_0(X1),u1_pre_topc(esk2_0))|~v1_tsep_1(X1,esk2_0)|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_103]), c_0_66])])).
fof(c_0_121, plain, ![X34, X35, X36]:(~v2_pre_topc(X34)|~l1_pre_topc(X34)|(~m1_pre_topc(X35,X34)|(~m1_pre_topc(X36,X35)|m1_pre_topc(X36,X34)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tsep_1])])])).
cnf(c_0_122, plain, (m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_123, plain, (m1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),X2)|~m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_87]), c_0_108])).
cnf(c_0_124, negated_conjecture, (m1_pre_topc(esk2_0,g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_66])])).
cnf(c_0_125, negated_conjecture, (l1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_54]), c_0_55])])).
fof(c_0_126, plain, ![X7]:(~l1_pre_topc(X7)|(~v1_pre_topc(X7)|X7=g1_pre_topc(u1_struct_0(X7),u1_pre_topc(X7)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_pre_topc])])).
fof(c_0_127, plain, ![X57, X58, X59]:((u1_struct_0(k8_tmap_1(X57,X58))=u1_struct_0(X57)|(v2_struct_0(X58)|~m1_pre_topc(X58,X57))|(v2_struct_0(X57)|(~v2_pre_topc(X57)|~l1_pre_topc(X57))))&(~m1_subset_1(X59,k1_zfmisc_1(u1_struct_0(X57)))|(X59!=u1_struct_0(X58)|u1_pre_topc(k8_tmap_1(X57,X58))=k5_tmap_1(X57,X59))|(v2_struct_0(X58)|~m1_pre_topc(X58,X57))|(v2_struct_0(X57)|(~v2_pre_topc(X57)|~l1_pre_topc(X57))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t114_tmap_1])])])])])).
cnf(c_0_128, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|~v1_xboole_0(u1_pre_topc(esk2_0))|~v1_tsep_1(esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_117, c_0_118])).
cnf(c_0_129, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),esk2_0)|v1_xboole_0(u1_pre_topc(esk2_0))|~v1_tsep_1(esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120]), c_0_66]), c_0_65])])).
cnf(c_0_130, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X1,X2)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_131, plain, (m1_pre_topc(X3,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_132, plain, (m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[c_0_122, c_0_52])).
cnf(c_0_133, negated_conjecture, (m1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124]), c_0_125]), c_0_66])])).
cnf(c_0_134, plain, (X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X1)|~v1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_126])).
cnf(c_0_135, plain, (u1_struct_0(k8_tmap_1(X1,X2))=u1_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_136, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),esk2_0)|v1_xboole_0(u1_struct_0(esk2_0))|~v1_tsep_1(esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_128, c_0_129])).
fof(c_0_137, plain, ![X45, X46, X47]:((v1_tsep_1(X46,X47)|~r1_tarski(u1_struct_0(X46),u1_struct_0(X47))|(v2_struct_0(X47)|~m1_pre_topc(X47,X45))|(~v1_tsep_1(X46,X45)|~m1_pre_topc(X46,X45))|(v2_struct_0(X45)|(~v2_pre_topc(X45)|~l1_pre_topc(X45))))&(m1_pre_topc(X46,X47)|~r1_tarski(u1_struct_0(X46),u1_struct_0(X47))|(v2_struct_0(X47)|~m1_pre_topc(X47,X45))|(~v1_tsep_1(X46,X45)|~m1_pre_topc(X46,X45))|(v2_struct_0(X45)|(~v2_pre_topc(X45)|~l1_pre_topc(X45))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t19_tsep_1])])])])])).
cnf(c_0_138, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X3)|~v2_pre_topc(X3)), inference(csr,[status(thm)],[c_0_130, c_0_131])).
cnf(c_0_139, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))=k8_tmap_1(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_140, plain, ![X64, X65]:((~m1_subset_1(X64,k1_zfmisc_1(X65))|r1_tarski(X64,X65))&(~r1_tarski(X64,X65)|m1_subset_1(X64,k1_zfmisc_1(X65)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_141, plain, (m1_pre_topc(X1,X2)|~m1_pre_topc(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3)),X2)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X2)|~l1_pre_topc(X3)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_131, c_0_132])).
cnf(c_0_142, negated_conjecture, (m1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_133]), c_0_66]), c_0_79])])).
fof(c_0_143, plain, ![X80, X81]:((v1_pre_topc(k6_tmap_1(X80,X81))|(v2_struct_0(X80)|(~v2_pre_topc(X80)|~l1_pre_topc(X80))|~m1_subset_1(X81,k1_zfmisc_1(u1_struct_0(X80)))))&((v2_pre_topc(k6_tmap_1(X80,X81))|(v2_struct_0(X80)|(~v2_pre_topc(X80)|~l1_pre_topc(X80))|~m1_subset_1(X81,k1_zfmisc_1(u1_struct_0(X80)))))&(l1_pre_topc(k6_tmap_1(X80,X81))|(v2_struct_0(X80)|(~v2_pre_topc(X80)|~l1_pre_topc(X80))|~m1_subset_1(X81,k1_zfmisc_1(u1_struct_0(X80))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_tmap_1])])])])).
fof(c_0_144, plain, ![X48]:(v2_struct_0(X48)|~l1_struct_0(X48)|~v1_xboole_0(u1_struct_0(X48))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_145, plain, ![X63]:(~l1_pre_topc(X63)|l1_struct_0(X63)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_146, plain, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(k8_tmap_1(X1,X2)))=k8_tmap_1(X1,X2)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_71]), c_0_72])).
cnf(c_0_147, negated_conjecture, (u1_pre_topc(k8_tmap_1(esk2_0,esk2_0))=u1_pre_topc(esk2_0)|v1_xboole_0(u1_struct_0(esk2_0))|~v1_tsep_1(esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_136]), c_0_65])])).
cnf(c_0_148, plain, (v1_tsep_1(X1,X2)|v2_struct_0(X2)|v2_struct_0(X3)|~r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X2,X3)|~v1_tsep_1(X1,X3)|~m1_pre_topc(X1,X3)|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_137])).
cnf(c_0_149, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk2_0))|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_54]), c_0_55]), c_0_56])])).
cnf(c_0_150, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|m1_pre_topc(X1,k8_tmap_1(esk1_0,esk2_0))|~m1_pre_topc(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_139]), c_0_55])])).
cnf(c_0_151, plain, (m1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_99]), c_0_108])).
cnf(c_0_152, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_153, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_140])).
cnf(c_0_154, negated_conjecture, (m1_pre_topc(X1,g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_142]), c_0_125]), c_0_66]), c_0_116])])).
cnf(c_0_155, plain, (l1_pre_topc(k6_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_143])).
cnf(c_0_156, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_140])).
cnf(c_0_157, plain, (u1_struct_0(k6_tmap_1(X1,X2))=u1_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_158, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_144])).
cnf(c_0_159, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_145])).
cnf(c_0_160, negated_conjecture, (g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))=k8_tmap_1(esk2_0,esk2_0)|v1_xboole_0(u1_struct_0(esk2_0))|~v1_tsep_1(esk2_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_147]), c_0_65]), c_0_66]), c_0_79])]), c_0_80])).
cnf(c_0_161, negated_conjecture, (v1_tsep_1(X1,esk2_0)|v2_struct_0(X2)|~v1_tsep_1(X1,X2)|~m1_pre_topc(esk2_0,X2)|~m1_pre_topc(X1,esk2_0)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_149]), c_0_80]), c_0_131])).
cnf(c_0_162, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_139]), c_0_55])])).
cnf(c_0_163, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_164, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|m1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),k8_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_151]), c_0_55])])).
cnf(c_0_165, negated_conjecture, (u1_struct_0(esk2_0)=u1_struct_0(X1)|~r1_tarski(u1_struct_0(esk2_0),u1_struct_0(X1))|~m1_pre_topc(X1,esk2_0)), inference(spm,[status(thm)],[c_0_152, c_0_149])).
cnf(c_0_166, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_153, c_0_51])).
cnf(c_0_167, negated_conjecture, (l1_pre_topc(X1)|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_154]), c_0_125])])).
cnf(c_0_168, plain, (l1_pre_topc(k6_tmap_1(X1,X2))|v2_struct_0(X1)|~r1_tarski(X2,u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_155, c_0_156])).
cnf(c_0_169, negated_conjecture, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_54]), c_0_55])])).
cnf(c_0_170, negated_conjecture, (u1_struct_0(k6_tmap_1(esk2_0,u1_struct_0(X1)))=u1_struct_0(esk2_0)|~m1_pre_topc(X1,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_78]), c_0_66]), c_0_79])]), c_0_80])).
cnf(c_0_171, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_158, c_0_159])).
cnf(c_0_172, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|m1_pre_topc(k8_tmap_1(esk2_0,esk2_0),esk2_0)|~v1_tsep_1(esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_160]), c_0_66])])).
cnf(c_0_173, negated_conjecture, (v1_tsep_1(esk2_0,esk2_0)|l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))|~m1_pre_topc(esk2_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_162]), c_0_54]), c_0_55]), c_0_56])]), c_0_163])).
cnf(c_0_174, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X1,k8_tmap_1(esk1_0,esk2_0))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_139]), c_0_55])])).
cnf(c_0_175, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|l1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_164]), c_0_162])).
cnf(c_0_176, negated_conjecture, (u1_struct_0(esk2_0)=u1_struct_0(X1)|~m1_pre_topc(X1,esk2_0)|~m1_pre_topc(esk2_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_166]), c_0_167])).
cnf(c_0_177, negated_conjecture, (l1_pre_topc(k6_tmap_1(esk2_0,u1_struct_0(X1)))|~m1_pre_topc(X1,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_149]), c_0_66]), c_0_79])]), c_0_80])).
cnf(c_0_178, negated_conjecture, (l1_pre_topc(k6_tmap_1(esk1_0,u1_struct_0(X1)))|~m1_pre_topc(X1,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_169]), c_0_55]), c_0_56])]), c_0_163])).
cnf(c_0_179, negated_conjecture, (u1_struct_0(k8_tmap_1(esk2_0,X1))=u1_struct_0(esk2_0)|~m1_pre_topc(X1,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_82]), c_0_66]), c_0_79])]), c_0_80])).
cnf(c_0_180, negated_conjecture, (m1_pre_topc(k8_tmap_1(esk2_0,esk2_0),esk2_0)|~v1_tsep_1(esk2_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_172]), c_0_66])]), c_0_80])).
cnf(c_0_181, negated_conjecture, (v1_tsep_1(esk2_0,esk2_0)|l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_173, c_0_65])])).
cnf(c_0_182, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|m1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_164]), c_0_175])).
fof(c_0_183, plain, ![X12]:(~v2_struct_0(X12)|~l1_struct_0(X12)|v1_xboole_0(u1_struct_0(X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_struct_0])])).
cnf(c_0_184, negated_conjecture, (u1_struct_0(k6_tmap_1(esk2_0,u1_struct_0(esk2_0)))=u1_struct_0(esk2_0)|~m1_pre_topc(X1,esk2_0)|~m1_pre_topc(esk2_0,X1)), inference(spm,[status(thm)],[c_0_170, c_0_176])).
cnf(c_0_185, negated_conjecture, (l1_pre_topc(k6_tmap_1(esk2_0,u1_struct_0(esk2_0)))|~m1_pre_topc(X1,esk2_0)|~m1_pre_topc(esk2_0,X1)), inference(spm,[status(thm)],[c_0_177, c_0_176])).
cnf(c_0_186, plain, (m1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_87]), c_0_52])).
cnf(c_0_187, plain, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(X1,k8_tmap_1(X2,X3))|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_135]), c_0_71])).
cnf(c_0_188, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|m1_pre_topc(k8_tmap_1(esk1_0,esk2_0),k8_tmap_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_164, c_0_139])).
cnf(c_0_189, negated_conjecture, (l1_pre_topc(k6_tmap_1(esk1_0,u1_struct_0(esk2_0)))|~m1_pre_topc(k8_tmap_1(esk2_0,X1),esk2_0)|~m1_pre_topc(X1,esk2_0)), inference(spm,[status(thm)],[c_0_178, c_0_179])).
cnf(c_0_190, negated_conjecture, (m1_pre_topc(k8_tmap_1(esk2_0,esk2_0),esk2_0)|l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_180, c_0_181])).
cnf(c_0_191, negated_conjecture, (m1_subset_1(u1_struct_0(X1),u1_pre_topc(esk1_0))|k5_tmap_1(esk1_0,u1_struct_0(X1))!=u1_pre_topc(esk1_0)|~m1_pre_topc(X1,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_169]), c_0_55]), c_0_56])]), c_0_163])).
cnf(c_0_192, negated_conjecture, (k5_tmap_1(esk1_0,u1_struct_0(X1))=u1_pre_topc(esk1_0)|~v3_pre_topc(u1_struct_0(X1),esk1_0)|~m1_pre_topc(X1,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_169]), c_0_55]), c_0_56])]), c_0_163])).
cnf(c_0_193, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|v2_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_182]), c_0_55]), c_0_56])])).
cnf(c_0_194, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_183])).
cnf(c_0_195, negated_conjecture, (u1_struct_0(k6_tmap_1(esk2_0,u1_struct_0(esk2_0)))=u1_struct_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_65]), c_0_65])])).
cnf(c_0_196, negated_conjecture, (l1_pre_topc(k6_tmap_1(esk2_0,u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_65]), c_0_65])])).
cnf(c_0_197, negated_conjecture, (m1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),esk1_0)|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_54]), c_0_55]), c_0_56])])).
cnf(c_0_198, plain, (v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_199, negated_conjecture, (m1_subset_1(u1_struct_0(k8_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|v1_tsep_1(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_188]), c_0_54]), c_0_55]), c_0_56])]), c_0_163]), c_0_80])).
cnf(c_0_200, negated_conjecture, (l1_pre_topc(k6_tmap_1(esk1_0,u1_struct_0(esk2_0)))|l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_190]), c_0_65])])).
cnf(c_0_201, negated_conjecture, (m1_subset_1(u1_struct_0(X1),u1_pre_topc(esk1_0))|~v3_pre_topc(u1_struct_0(X1),esk1_0)|~m1_pre_topc(X1,esk2_0)), inference(spm,[status(thm)],[c_0_191, c_0_192])).
cnf(c_0_202, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|m1_pre_topc(esk1_0,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_182]), c_0_55])]), c_0_193])).
cnf(c_0_203, plain, (v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_struct_0(X1)), inference(spm,[status(thm)],[c_0_194, c_0_159])).
cnf(c_0_204, negated_conjecture, (u1_struct_0(k8_tmap_1(esk2_0,esk2_0))=u1_struct_0(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_82]), c_0_65]), c_0_66]), c_0_79])]), c_0_80])).
cnf(c_0_205, negated_conjecture, (l1_pre_topc(k8_tmap_1(esk2_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_82]), c_0_65]), c_0_66]), c_0_79])]), c_0_80])).
cnf(c_0_206, negated_conjecture, (m1_pre_topc(X1,esk1_0)|~v1_pre_topc(X1)|~m1_pre_topc(X1,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_134]), c_0_167])).
cnf(c_0_207, negated_conjecture, (v1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_54]), c_0_55])])).
cnf(c_0_208, plain, (u1_struct_0(X1)=u1_struct_0(X2)|~r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_152, c_0_166])).
cnf(c_0_209, negated_conjecture, (r1_tarski(u1_struct_0(k8_tmap_1(esk1_0,esk2_0)),u1_struct_0(esk1_0))|v1_tsep_1(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_153, c_0_199])).
cnf(c_0_210, negated_conjecture, (l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_82]), c_0_54]), c_0_55]), c_0_56])]), c_0_163])).
cnf(c_0_211, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|m1_pre_topc(esk1_0,k8_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_99]), c_0_55])])).
cnf(c_0_212, negated_conjecture, (m1_subset_1(u1_struct_0(esk2_0),u1_pre_topc(esk1_0))|~v3_pre_topc(u1_struct_0(esk2_0),esk1_0)|~m1_pre_topc(k8_tmap_1(esk2_0,X1),esk2_0)|~m1_pre_topc(X1,esk2_0)), inference(spm,[status(thm)],[c_0_201, c_0_179])).
cnf(c_0_213, negated_conjecture, (v1_tsep_1(esk2_0,esk2_0)|m1_pre_topc(esk1_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_202]), c_0_54]), c_0_65]), c_0_55]), c_0_56])]), c_0_163])).
cnf(c_0_214, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|~v2_struct_0(k8_tmap_1(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_204]), c_0_205])])).
cnf(c_0_215, negated_conjecture, (m1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_207]), c_0_133])])).
cnf(c_0_216, plain, (v1_tsep_1(X1,X1)|v2_struct_0(X1)|v2_struct_0(X2)|~v1_tsep_1(X1,X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_148, c_0_43])).
cnf(c_0_217, negated_conjecture, (u1_struct_0(k8_tmap_1(esk1_0,esk2_0))=u1_struct_0(esk1_0)|v1_tsep_1(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_209]), c_0_210])]), c_0_211])).
cnf(c_0_218, negated_conjecture, (m1_subset_1(u1_struct_0(esk2_0),u1_pre_topc(esk1_0))|~v1_tsep_1(esk2_0,esk1_0)|~m1_pre_topc(k8_tmap_1(esk2_0,X1),esk2_0)|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_103]), c_0_54]), c_0_55])])).
cnf(c_0_219, negated_conjecture, (m1_pre_topc(k8_tmap_1(esk2_0,esk2_0),esk2_0)|m1_pre_topc(esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_180, c_0_213])).
cnf(c_0_220, negated_conjecture, (v2_struct_0(k6_tmap_1(esk2_0,u1_struct_0(esk2_0)))|~v1_xboole_0(u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_195]), c_0_196])])).
cnf(c_0_221, negated_conjecture, (~v2_struct_0(k8_tmap_1(esk2_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_214]), c_0_66])]), c_0_80])).
cnf(c_0_222, plain, (v3_pre_topc(u1_struct_0(X1),X2)|v1_xboole_0(u1_pre_topc(X2))|~m1_subset_1(u1_struct_0(X1),u1_pre_topc(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_112, c_0_51])).
cnf(c_0_223, negated_conjecture, (m1_subset_1(u1_struct_0(X1),u1_pre_topc(esk1_0))|~v1_tsep_1(X1,esk1_0)|~m1_pre_topc(X1,esk2_0)|~m1_pre_topc(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_103]), c_0_55])])).
cnf(c_0_224, negated_conjecture, (m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_215]), c_0_55]), c_0_66]), c_0_56])])).
cnf(c_0_225, negated_conjecture, (u1_struct_0(k6_tmap_1(esk1_0,u1_struct_0(X1)))=u1_struct_0(esk1_0)|~m1_pre_topc(X1,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_169]), c_0_55]), c_0_56])]), c_0_163])).
cnf(c_0_226, negated_conjecture, (u1_struct_0(k8_tmap_1(esk1_0,esk2_0))=u1_struct_0(esk1_0)|v1_tsep_1(esk2_0,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_217]), c_0_54]), c_0_55]), c_0_56])]), c_0_80]), c_0_163])).
cnf(c_0_227, negated_conjecture, (m1_subset_1(u1_struct_0(esk2_0),u1_pre_topc(esk1_0))|m1_pre_topc(esk1_0,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_219]), c_0_65])]), c_0_202])).
cnf(c_0_228, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_82]), c_0_65]), c_0_66]), c_0_79])]), c_0_221]), c_0_80])).
cnf(c_0_229, negated_conjecture, (v3_pre_topc(u1_struct_0(X1),esk1_0)|v1_xboole_0(u1_pre_topc(esk1_0))|~v1_tsep_1(X1,esk1_0)|~m1_pre_topc(X1,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_223]), c_0_55])]), c_0_224])).
cnf(c_0_230, negated_conjecture, (u1_struct_0(k6_tmap_1(esk1_0,u1_struct_0(esk2_0)))=u1_struct_0(esk1_0)|~m1_pre_topc(k8_tmap_1(esk2_0,esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_225, c_0_204])).
cnf(c_0_231, negated_conjecture, (u1_struct_0(k8_tmap_1(esk1_0,esk2_0))=u1_struct_0(esk1_0)|m1_pre_topc(k8_tmap_1(esk2_0,esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_180, c_0_226])).
fof(c_0_232, plain, ![X11]:((v1_pre_topc(g1_pre_topc(u1_struct_0(X11),u1_pre_topc(X11)))|(~v2_pre_topc(X11)|~l1_pre_topc(X11)))&(v2_pre_topc(g1_pre_topc(u1_struct_0(X11),u1_pre_topc(X11)))|(~v2_pre_topc(X11)|~l1_pre_topc(X11)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc6_pre_topc])])])).
cnf(c_0_233, negated_conjecture, (m1_pre_topc(esk1_0,esk1_0)|~v1_xboole_0(u1_pre_topc(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_227]), c_0_228])).
cnf(c_0_234, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),esk1_0)|v1_xboole_0(u1_pre_topc(esk1_0))|m1_pre_topc(esk1_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_202]), c_0_65])])).
cnf(c_0_235, negated_conjecture, (u1_struct_0(k8_tmap_1(esk1_0,esk2_0))=u1_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_82]), c_0_54]), c_0_55]), c_0_56])]), c_0_163]), c_0_231])).
cnf(c_0_236, plain, (v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_232])).
cnf(c_0_237, negated_conjecture, (u1_pre_topc(k8_tmap_1(esk1_0,X1))=u1_pre_topc(esk1_0)|~v3_pre_topc(u1_struct_0(X1),esk1_0)|~m1_pre_topc(X1,esk2_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_91]), c_0_55]), c_0_56])]), c_0_163]), c_0_224])).
cnf(c_0_238, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),esk1_0)|m1_pre_topc(esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_233, c_0_234])).
cnf(c_0_239, negated_conjecture, (v2_pre_topc(k8_tmap_1(esk1_0,esk2_0))|~v2_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(k8_tmap_1(esk1_0,esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_235]), c_0_210])])).
cnf(c_0_240, plain, (v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(k8_tmap_1(X1,X2))))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_135]), c_0_70]), c_0_71])).
cnf(c_0_241, negated_conjecture, (u1_pre_topc(k8_tmap_1(esk1_0,esk2_0))=u1_pre_topc(esk1_0)|m1_pre_topc(esk1_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_238]), c_0_65])])).
cnf(c_0_242, negated_conjecture, (v2_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_240]), c_0_54]), c_0_55]), c_0_56])]), c_0_163]), c_0_80])).
cnf(c_0_243, plain, (m1_pre_topc(X1,X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_132]), c_0_108])).
cnf(c_0_244, negated_conjecture, (m1_pre_topc(esk1_0,esk1_0)|v2_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_241]), c_0_235]), c_0_210]), c_0_242])])).
cnf(c_0_245, negated_conjecture, (m1_pre_topc(esk1_0,esk1_0)|m1_pre_topc(X1,X1)|~m1_pre_topc(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_244]), c_0_55])])).
cnf(c_0_246, negated_conjecture, (m1_pre_topc(esk1_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_99]), c_0_55])])).
cnf(c_0_247, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X1,X2)|~m1_pre_topc(X2,X3)|~l1_pre_topc(X3)|~v2_pre_topc(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_132]), c_0_108])).
cnf(c_0_248, negated_conjecture, (v2_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_246]), c_0_55]), c_0_56])])).
cnf(c_0_249, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_248]), c_0_55])])).
fof(c_0_250, plain, ![X31, X32, X33]:(v2_struct_0(X31)|(~v2_pre_topc(X31)|~l1_pre_topc(X31))|(~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|(~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(k6_tmap_1(X31,X32))))|(X33!=X32|v3_pre_topc(X33,k6_tmap_1(X31,X32)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t105_tmap_1])])])])).
cnf(c_0_251, negated_conjecture, (u1_struct_0(X1)=u1_struct_0(X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_208, c_0_249])).
cnf(c_0_252, plain, (v2_struct_0(X1)|v3_pre_topc(X3,k6_tmap_1(X1,X2))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k6_tmap_1(X1,X2))))|X3!=X2), inference(split_conjunct,[status(thm)],[c_0_250])).
cnf(c_0_253, negated_conjecture, (u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))=u1_struct_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_124]), c_0_54]), c_0_133]), c_0_125])])).
cnf(c_0_254, plain, (v3_pre_topc(X1,k6_tmap_1(X2,X1))|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k6_tmap_1(X2,X1))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(er,[status(thm)],[c_0_252])).
cnf(c_0_255, negated_conjecture, (m1_pre_topc(esk2_0,X1)|~m1_pre_topc(esk1_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_131, c_0_54])).
cnf(c_0_256, negated_conjecture, (k5_tmap_1(esk1_0,u1_struct_0(esk2_0))=u1_pre_topc(esk1_0)|~v3_pre_topc(u1_struct_0(esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_253]), c_0_133])])).
cnf(c_0_257, plain, (v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_198, c_0_99])).
cnf(c_0_258, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~m1_pre_topc(X1,esk2_0)|~v2_struct_0(k6_tmap_1(esk1_0,u1_struct_0(X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_225]), c_0_178])).
cnf(c_0_259, negated_conjecture, (v3_pre_topc(u1_struct_0(X1),k6_tmap_1(esk1_0,u1_struct_0(X1)))|~m1_pre_topc(X1,esk2_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_225]), c_0_55]), c_0_56])]), c_0_163]), c_0_169])).
cnf(c_0_260, negated_conjecture, (m1_pre_topc(esk2_0,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_pre_topc(esk1_0,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_236]), c_0_108])).
cnf(c_0_261, negated_conjecture, (u1_pre_topc(k8_tmap_1(esk1_0,esk2_0))=u1_pre_topc(esk1_0)|~v3_pre_topc(u1_struct_0(esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_256]), c_0_54]), c_0_55]), c_0_56])]), c_0_163])).
cnf(c_0_262, negated_conjecture, (v1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(k8_tmap_1(esk1_0,esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_235]), c_0_210])])).
cnf(c_0_263, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~m1_pre_topc(X1,esk2_0)|~m1_pre_topc(esk2_0,X1)|~v2_struct_0(k6_tmap_1(esk1_0,u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_258, c_0_176])).
cnf(c_0_264, plain, (v1_tsep_1(X2,X1)|~v3_pre_topc(esk3_2(X1,X2),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_265, plain, (esk3_2(X1,X2)=u1_struct_0(X2)|v1_tsep_1(X2,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_266, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),k6_tmap_1(esk1_0,u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_253]), c_0_133])])).
cnf(c_0_267, negated_conjecture, (m1_pre_topc(esk2_0,X1)|~m1_pre_topc(esk1_0,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_260]), c_0_66])])).
cnf(c_0_268, negated_conjecture, (u1_pre_topc(k8_tmap_1(esk1_0,esk2_0))=u1_pre_topc(esk1_0)|~v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_103]), c_0_54]), c_0_55])])).
cnf(c_0_269, negated_conjecture, (m1_pre_topc(esk1_0,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_248]), c_0_55])])).
cnf(c_0_270, negated_conjecture, (v1_tsep_1(esk2_0,esk1_0)|m1_pre_topc(esk2_0,k8_tmap_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_150, c_0_54])).
cnf(c_0_271, negated_conjecture, (v1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_146]), c_0_54]), c_0_55]), c_0_56])]), c_0_163]), c_0_80])).
cnf(c_0_272, negated_conjecture, (~v1_tsep_1(esk2_0,esk1_0)|~m1_pre_topc(esk2_0,esk1_0)|g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))!=k8_tmap_1(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_273, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~m1_pre_topc(X1,esk2_0)|~m1_pre_topc(esk2_0,X1)|~v2_struct_0(k8_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_82]), c_0_54]), c_0_55]), c_0_56])]), c_0_163])).
cnf(c_0_274, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk1_0))|~m1_pre_topc(X1,esk2_0)), inference(spm,[status(thm)],[c_0_153, c_0_169])).
cnf(c_0_275, plain, (v1_tsep_1(X1,X2)|~v3_pre_topc(u1_struct_0(X1),X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_264, c_0_265])).
cnf(c_0_276, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),k8_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_82]), c_0_54]), c_0_55]), c_0_56])]), c_0_163])).
cnf(c_0_277, negated_conjecture, (m1_pre_topc(esk2_0,k8_tmap_1(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_268]), c_0_235]), c_0_269]), c_0_210]), c_0_242])]), c_0_270])).
cnf(c_0_278, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))=k8_tmap_1(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_268]), c_0_235]), c_0_271]), c_0_210])]), c_0_139])).
cnf(c_0_279, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))!=k8_tmap_1(esk1_0,esk2_0)|~v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_272, c_0_54])])).
cnf(c_0_280, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~v2_struct_0(k8_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_65]), c_0_65])])).
cnf(c_0_281, negated_conjecture, (v1_tsep_1(X1,esk1_0)|v2_struct_0(X2)|~v1_tsep_1(X1,X2)|~m1_pre_topc(esk1_0,X2)|~m1_pre_topc(X1,esk2_0)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_274]), c_0_163])).
cnf(c_0_282, negated_conjecture, (v1_tsep_1(esk2_0,k8_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_276]), c_0_277]), c_0_210])])).
cnf(c_0_283, negated_conjecture, (m1_pre_topc(esk1_0,k8_tmap_1(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_269, c_0_278])).
cnf(c_0_284, negated_conjecture, (~v1_tsep_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_279, c_0_278])])).
cnf(c_0_285, negated_conjecture, (~v2_struct_0(k8_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_280]), c_0_55])]), c_0_163])).
cnf(c_0_286, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_281, c_0_282]), c_0_283]), c_0_65]), c_0_277]), c_0_210]), c_0_242])]), c_0_284]), c_0_285]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 287
# Proof object clause steps            : 220
# Proof object formula steps           : 67
# Proof object conjectures             : 137
# Proof object clause conjectures      : 134
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 52
# Proof object initial formulas used   : 33
# Proof object generating inferences   : 156
# Proof object simplifying inferences  : 356
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 38
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 76
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 75
# Processed clauses                    : 16920
# ...of these trivial                  : 505
# ...subsumed                          : 11979
# ...remaining for further processing  : 4436
# Other redundant clauses eliminated   : 10
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1085
# Backward-rewritten                   : 492
# Generated clauses                    : 217212
# ...of the previous two non-trivial   : 195247
# Contextual simplify-reflections      : 1028
# Paramodulations                      : 217149
# Factorizations                       : 0
# Equation resolutions                 : 11
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2730
#    Positive orientable unit clauses  : 307
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 62
#    Non-unit-clauses                  : 2361
# Current number of unprocessed clauses: 174827
# ...number of literals in the above   : 1131840
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1699
# Clause-clause subsumption calls (NU) : 1504747
# Rec. Clause-clause subsumption calls : 440552
# Non-unit clause-clause subsumptions  : 10218
# Unit Clause-clause subsumption calls : 61758
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3082
# BW rewrite match successes           : 118
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6614219

# -------------------------------------------------
# User time                : 4.118 s
# System time              : 0.053 s
# Total time               : 4.171 s
# Maximum resident set size: 147284 pages
