# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t204_relat_1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)<=>r2_hidden(X2,k11_relat_1(X3,X1)))), file('data/MPTP/proofs/5258bd78f4d8453185b7319fbe996ebf.E_input', t204_relat_1)).
fof(d16_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:k11_relat_1(X1,X2)=k9_relat_1(X1,k1_tarski(X2))), file('data/MPTP/proofs/5258bd78f4d8453185b7319fbe996ebf.E_input', d16_relat_1)).
fof(d13_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(X3=k9_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(k4_tarski(X5,X4),X1)&r2_hidden(X5,X2))))), file('data/MPTP/proofs/5258bd78f4d8453185b7319fbe996ebf.E_input', d13_relat_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/5258bd78f4d8453185b7319fbe996ebf.E_input', d1_tarski)).
fof(t143_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k9_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k1_relat_1(X3))&(r2_hidden(k4_tarski(X4,X1),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/5258bd78f4d8453185b7319fbe996ebf.E_input', t143_relat_1)).
fof(c_0_5, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)<=>r2_hidden(X2,k11_relat_1(X3,X1))))), inference(assume_negation,[status(cth)],[t204_relat_1])).
fof(c_0_6, plain, ![X9, X10]:(~v1_relat_1(X9)|k11_relat_1(X9,X10)=k9_relat_1(X9,k1_tarski(X10))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d16_relat_1])])])).
fof(c_0_7, negated_conjecture, (v1_relat_1(esk3_0)&((~r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|~r2_hidden(esk2_0,k11_relat_1(esk3_0,esk1_0)))&(r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r2_hidden(esk2_0,k11_relat_1(esk3_0,esk1_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
cnf(c_0_8, plain, (k11_relat_1(X1,X2)=k9_relat_1(X1,k1_tarski(X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_10, plain, ![X11, X12, X13, X14, X16, X17, X18, X19, X21]:((((r2_hidden(k4_tarski(esk4_4(X11,X12,X13,X14),X14),X11)|~r2_hidden(X14,X13)|X13!=k9_relat_1(X11,X12)|~v1_relat_1(X11))&(r2_hidden(esk4_4(X11,X12,X13,X14),X12)|~r2_hidden(X14,X13)|X13!=k9_relat_1(X11,X12)|~v1_relat_1(X11)))&(~r2_hidden(k4_tarski(X17,X16),X11)|~r2_hidden(X17,X12)|r2_hidden(X16,X13)|X13!=k9_relat_1(X11,X12)|~v1_relat_1(X11)))&((~r2_hidden(esk5_3(X11,X18,X19),X19)|(~r2_hidden(k4_tarski(X21,esk5_3(X11,X18,X19)),X11)|~r2_hidden(X21,X18))|X19=k9_relat_1(X11,X18)|~v1_relat_1(X11))&((r2_hidden(k4_tarski(esk6_3(X11,X18,X19),esk5_3(X11,X18,X19)),X11)|r2_hidden(esk5_3(X11,X18,X19),X19)|X19=k9_relat_1(X11,X18)|~v1_relat_1(X11))&(r2_hidden(esk6_3(X11,X18,X19),X18)|r2_hidden(esk5_3(X11,X18,X19),X19)|X19=k9_relat_1(X11,X18)|~v1_relat_1(X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d13_relat_1])])])])])])).
cnf(c_0_11, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r2_hidden(esk2_0,k11_relat_1(esk3_0,esk1_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, negated_conjecture, (k11_relat_1(esk3_0,X1)=k9_relat_1(esk3_0,k1_tarski(X1))), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_13, plain, (r2_hidden(X2,X5)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,X4)|X5!=k9_relat_1(X3,X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, negated_conjecture, (r2_hidden(esk2_0,k9_relat_1(esk3_0,k1_tarski(esk1_0)))|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(rw,[status(thm)],[c_0_11, c_0_12])).
fof(c_0_15, plain, ![X23, X24, X25, X26, X27, X28]:(((~r2_hidden(X25,X24)|X25=X23|X24!=k1_tarski(X23))&(X26!=X23|r2_hidden(X26,X24)|X24!=k1_tarski(X23)))&((~r2_hidden(esk7_2(X27,X28),X28)|esk7_2(X27,X28)!=X27|X28=k1_tarski(X27))&(r2_hidden(esk7_2(X27,X28),X28)|esk7_2(X27,X28)=X27|X28=k1_tarski(X27)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_16, negated_conjecture, (r2_hidden(esk2_0,k9_relat_1(esk3_0,k1_tarski(esk1_0)))|r2_hidden(esk2_0,X1)|X1!=k9_relat_1(esk3_0,X2)|~r2_hidden(esk1_0,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_9])])).
cnf(c_0_17, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_18, plain, ![X30, X31, X32, X34]:(((r2_hidden(esk8_3(X30,X31,X32),k1_relat_1(X32))|~r2_hidden(X30,k9_relat_1(X32,X31))|~v1_relat_1(X32))&((r2_hidden(k4_tarski(esk8_3(X30,X31,X32),X30),X32)|~r2_hidden(X30,k9_relat_1(X32,X31))|~v1_relat_1(X32))&(r2_hidden(esk8_3(X30,X31,X32),X31)|~r2_hidden(X30,k9_relat_1(X32,X31))|~v1_relat_1(X32))))&(~r2_hidden(X34,k1_relat_1(X32))|(~r2_hidden(k4_tarski(X34,X30),X32)|~r2_hidden(X34,X31))|r2_hidden(X30,k9_relat_1(X32,X31))|~v1_relat_1(X32))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t143_relat_1])])])])])).
cnf(c_0_19, negated_conjecture, (r2_hidden(esk2_0,k9_relat_1(esk3_0,k1_tarski(esk1_0)))|r2_hidden(esk2_0,k9_relat_1(esk3_0,X1))|~r2_hidden(esk1_0,X1)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, (r2_hidden(X1,X2)|X2!=k1_tarski(X1)), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_21, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (r2_hidden(esk8_3(X1,X2,X3),X2)|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, negated_conjecture, (r2_hidden(esk2_0,k9_relat_1(esk3_0,k1_tarski(esk1_0)))|r2_hidden(esk2_0,k9_relat_1(esk3_0,X1))|X1!=k1_tarski(esk1_0)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_24, plain, (esk8_3(X1,X2,X3)=X4|X2!=k1_tarski(X4)|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_25, negated_conjecture, (r2_hidden(esk2_0,k9_relat_1(esk3_0,k1_tarski(esk1_0)))), inference(er,[status(thm)],[c_0_23])).
cnf(c_0_26, negated_conjecture, (~r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|~r2_hidden(esk2_0,k11_relat_1(esk3_0,esk1_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_27, negated_conjecture, (esk8_3(esk2_0,k1_tarski(esk1_0),esk3_0)=X1|k1_tarski(esk1_0)!=k1_tarski(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_9])])).
cnf(c_0_28, negated_conjecture, (~r2_hidden(esk2_0,k9_relat_1(esk3_0,k1_tarski(esk1_0)))|~r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(rw,[status(thm)],[c_0_26, c_0_12])).
cnf(c_0_29, plain, (r2_hidden(k4_tarski(esk8_3(X1,X2,X3),X1),X3)|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_30, negated_conjecture, (esk8_3(esk2_0,k1_tarski(esk1_0),esk3_0)=esk1_0), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_31, negated_conjecture, (~r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_25])])).
cnf(c_0_32, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_25]), c_0_9])]), c_0_31]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 33
# Proof object clause steps            : 22
# Proof object formula steps           : 11
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 13
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 18
# Processed clauses                    : 61
# ...of these trivial                  : 1
# ...subsumed                          : 2
# ...remaining for further processing  : 58
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 12
# Generated clauses                    : 55
# ...of the previous two non-trivial   : 55
# Contextual simplify-reflections      : 0
# Paramodulations                      : 49
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 27
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 21
# Current number of unprocessed clauses: 27
# ...number of literals in the above   : 132
# Current number of archived formulas  : 0
# Current number of archived clauses   : 30
# Clause-clause subsumption calls (NU) : 439
# Rec. Clause-clause subsumption calls : 211
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2258

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28584 pages
