# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PI_S0a
# and selection function SelectNegativeLiterals.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t37_yellow19, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(v1_compts_1(X1)<=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>?[X3]:(m2_yellow_6(X3,X1,X2)&v3_yellow_6(X3,X1))))), file('data/MPTP/proofs/52c0d80ec41743fcb331cae698840740.E_input', t37_yellow19)).
fof(t35_yellow19, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(v1_compts_1(X1)<=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&r3_waybel_9(X1,X2,X3))))), file('data/MPTP/proofs/52c0d80ec41743fcb331cae698840740.E_input', t35_yellow19)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/52c0d80ec41743fcb331cae698840740.E_input', dt_l1_pre_topc)).
fof(dt_m2_yellow_6, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&(~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1)))))=>![X3]:(m2_yellow_6(X3,X1,X2)=>(~(v2_struct_0(X3))&(v4_orders_2(X3)&(v7_waybel_0(X3)&l1_waybel_0(X3,X1)))))), file('data/MPTP/proofs/52c0d80ec41743fcb331cae698840740.E_input', dt_m2_yellow_6)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/52c0d80ec41743fcb331cae698840740.E_input', t4_subset)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/52c0d80ec41743fcb331cae698840740.E_input', d3_tarski)).
fof(dt_k10_yellow_6, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&(~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1)))))=>m1_subset_1(k10_yellow_6(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/52c0d80ec41743fcb331cae698840740.E_input', dt_k10_yellow_6)).
fof(t29_waybel_9, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k10_yellow_6(X1,X2))=>r3_waybel_9(X1,X2,X3))))), file('data/MPTP/proofs/52c0d80ec41743fcb331cae698840740.E_input', t29_waybel_9)).
fof(t31_waybel_9, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m2_yellow_6(X3,X1,X2)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_waybel_9(X1,X3,X4)=>r3_waybel_9(X1,X2,X4)))))), file('data/MPTP/proofs/52c0d80ec41743fcb331cae698840740.E_input', t31_waybel_9)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/52c0d80ec41743fcb331cae698840740.E_input', d10_xboole_0)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/52c0d80ec41743fcb331cae698840740.E_input', t2_xboole_1)).
fof(d19_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>(v3_yellow_6(X2,X1)<=>~(k10_yellow_6(X1,X2)=k1_xboole_0)))), file('data/MPTP/proofs/52c0d80ec41743fcb331cae698840740.E_input', d19_yellow_6)).
fof(t32_waybel_9, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~((r3_waybel_9(X1,X2,X3)&![X4]:(m2_yellow_6(X4,X1,X2)=>~(r2_hidden(X3,k10_yellow_6(X1,X4))))))))), file('data/MPTP/proofs/52c0d80ec41743fcb331cae698840740.E_input', t32_waybel_9)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/52c0d80ec41743fcb331cae698840740.E_input', t7_ordinal1)).
fof(c_0_14, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(v1_compts_1(X1)<=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>?[X3]:(m2_yellow_6(X3,X1,X2)&v3_yellow_6(X3,X1)))))), inference(assume_negation,[status(cth)],[t37_yellow19])).
fof(c_0_15, plain, ![X10, X11, X14]:(((m1_subset_1(esk4_2(X10,X11),u1_struct_0(X10))|(v2_struct_0(X11)|(~v4_orders_2(X11)|(~v7_waybel_0(X11)|~l1_waybel_0(X11,X10))))|~v1_compts_1(X10)|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))&(r3_waybel_9(X10,X11,esk4_2(X10,X11))|(v2_struct_0(X11)|(~v4_orders_2(X11)|(~v7_waybel_0(X11)|~l1_waybel_0(X11,X10))))|~v1_compts_1(X10)|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10)))))&(((~v2_struct_0(esk5_1(X10))|v1_compts_1(X10)|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))&((v4_orders_2(esk5_1(X10))|v1_compts_1(X10)|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))&((v7_waybel_0(esk5_1(X10))|v1_compts_1(X10)|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))&(l1_waybel_0(esk5_1(X10),X10)|v1_compts_1(X10)|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10)))))))&(~m1_subset_1(X14,u1_struct_0(X10))|~r3_waybel_9(X10,esk5_1(X10),X14)|v1_compts_1(X10)|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t35_yellow19])])])])])])).
fof(c_0_16, negated_conjecture, ![X7, X8]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((((~v2_struct_0(esk2_0)|~v1_compts_1(esk1_0))&((v4_orders_2(esk2_0)|~v1_compts_1(esk1_0))&((v7_waybel_0(esk2_0)|~v1_compts_1(esk1_0))&(l1_waybel_0(esk2_0,esk1_0)|~v1_compts_1(esk1_0)))))&(~m2_yellow_6(X7,esk1_0,esk2_0)|~v3_yellow_6(X7,esk1_0)|~v1_compts_1(esk1_0)))&((m2_yellow_6(esk3_1(X8),esk1_0,X8)|(v2_struct_0(X8)|(~v4_orders_2(X8)|(~v7_waybel_0(X8)|~l1_waybel_0(X8,esk1_0))))|v1_compts_1(esk1_0))&(v3_yellow_6(esk3_1(X8),esk1_0)|(v2_struct_0(X8)|(~v4_orders_2(X8)|(~v7_waybel_0(X8)|~l1_waybel_0(X8,esk1_0))))|v1_compts_1(esk1_0))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])])])).
cnf(c_0_17, plain, (l1_waybel_0(esk5_1(X1),X1)|v1_compts_1(X1)|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_18, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_19, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, plain, (v4_orders_2(esk5_1(X1))|v1_compts_1(X1)|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (v7_waybel_0(esk5_1(X1))|v1_compts_1(X1)|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_23, plain, ![X15]:(~l1_pre_topc(X15)|l1_struct_0(X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_24, plain, ![X16, X17, X18]:((~v2_struct_0(X18)|~m2_yellow_6(X18,X16,X17)|(v2_struct_0(X16)|~l1_struct_0(X16)|(v2_struct_0(X17)|(~v4_orders_2(X17)|(~v7_waybel_0(X17)|~l1_waybel_0(X17,X16))))))&((v4_orders_2(X18)|~m2_yellow_6(X18,X16,X17)|(v2_struct_0(X16)|~l1_struct_0(X16)|(v2_struct_0(X17)|(~v4_orders_2(X17)|(~v7_waybel_0(X17)|~l1_waybel_0(X17,X16))))))&((v7_waybel_0(X18)|~m2_yellow_6(X18,X16,X17)|(v2_struct_0(X16)|~l1_struct_0(X16)|(v2_struct_0(X17)|(~v4_orders_2(X17)|(~v7_waybel_0(X17)|~l1_waybel_0(X17,X16))))))&(l1_waybel_0(X18,X16)|~m2_yellow_6(X18,X16,X17)|(v2_struct_0(X16)|~l1_struct_0(X16)|(v2_struct_0(X17)|(~v4_orders_2(X17)|(~v7_waybel_0(X17)|~l1_waybel_0(X17,X16))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_yellow_6])])])])])).
cnf(c_0_25, negated_conjecture, (m2_yellow_6(esk3_1(X1),esk1_0,X1)|v2_struct_0(X1)|v1_compts_1(esk1_0)|~v4_orders_2(X1)|~v7_waybel_0(X1)|~l1_waybel_0(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26, negated_conjecture, (l1_waybel_0(esk5_1(esk1_0),esk1_0)|v1_compts_1(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_27, negated_conjecture, (v4_orders_2(esk5_1(esk1_0))|v1_compts_1(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_28, negated_conjecture, (v7_waybel_0(esk5_1(esk1_0))|v1_compts_1(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_29, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_30, plain, ![X19, X20, X21]:(~r2_hidden(X19,X20)|~m1_subset_1(X20,k1_zfmisc_1(X21))|m1_subset_1(X19,X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_31, plain, ![X22, X23, X24, X25, X26]:((~r1_tarski(X22,X23)|(~r2_hidden(X24,X22)|r2_hidden(X24,X23)))&((r2_hidden(esk6_2(X25,X26),X25)|r1_tarski(X25,X26))&(~r2_hidden(esk6_2(X25,X26),X26)|r1_tarski(X25,X26)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_32, plain, ![X28, X29]:(v2_struct_0(X28)|(~v2_pre_topc(X28)|~l1_pre_topc(X28))|(v2_struct_0(X29)|(~v4_orders_2(X29)|(~v7_waybel_0(X29)|~l1_waybel_0(X29,X28))))|m1_subset_1(k10_yellow_6(X28,X29),k1_zfmisc_1(u1_struct_0(X28)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k10_yellow_6])])])).
cnf(c_0_33, plain, (l1_waybel_0(X1,X2)|v2_struct_0(X2)|v2_struct_0(X3)|~m2_yellow_6(X1,X2,X3)|~l1_struct_0(X2)|~v4_orders_2(X3)|~v7_waybel_0(X3)|~l1_waybel_0(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, negated_conjecture, (m2_yellow_6(esk3_1(esk5_1(esk1_0)),esk1_0,esk5_1(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(esk5_1(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_28])).
cnf(c_0_35, negated_conjecture, (l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_29, c_0_19])).
cnf(c_0_36, plain, (v4_orders_2(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m2_yellow_6(X1,X2,X3)|~l1_struct_0(X2)|~v4_orders_2(X3)|~v7_waybel_0(X3)|~l1_waybel_0(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_37, plain, (v7_waybel_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m2_yellow_6(X1,X2,X3)|~l1_struct_0(X2)|~v4_orders_2(X3)|~v7_waybel_0(X3)|~l1_waybel_0(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_38, plain, (v2_struct_0(X2)|v2_struct_0(X3)|~v2_struct_0(X1)|~m2_yellow_6(X1,X2,X3)|~l1_struct_0(X2)|~v4_orders_2(X3)|~v7_waybel_0(X3)|~l1_waybel_0(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_39, plain, ![X30, X31, X32]:(v2_struct_0(X30)|(~v2_pre_topc(X30)|~l1_pre_topc(X30))|(v2_struct_0(X31)|(~v4_orders_2(X31)|(~v7_waybel_0(X31)|~l1_waybel_0(X31,X30)))|(~m1_subset_1(X32,u1_struct_0(X30))|(~r2_hidden(X32,k10_yellow_6(X30,X31))|r3_waybel_9(X30,X31,X32))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t29_waybel_9])])])])).
cnf(c_0_40, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_41, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_42, plain, (v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(k10_yellow_6(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43, negated_conjecture, (l1_waybel_0(esk3_1(esk5_1(esk1_0)),esk1_0)|v1_compts_1(esk1_0)|v2_struct_0(esk5_1(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])]), c_0_20]), c_0_27]), c_0_28]), c_0_26])).
cnf(c_0_44, negated_conjecture, (v4_orders_2(esk3_1(esk5_1(esk1_0)))|v1_compts_1(esk1_0)|v2_struct_0(esk5_1(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_34]), c_0_35])]), c_0_20]), c_0_27]), c_0_28]), c_0_26])).
cnf(c_0_45, negated_conjecture, (v7_waybel_0(esk3_1(esk5_1(esk1_0)))|v1_compts_1(esk1_0)|v2_struct_0(esk5_1(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_34]), c_0_35])]), c_0_20]), c_0_27]), c_0_28]), c_0_26])).
cnf(c_0_46, negated_conjecture, (v1_compts_1(esk1_0)|v2_struct_0(esk5_1(esk1_0))|~v2_struct_0(esk3_1(esk5_1(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_34]), c_0_35])]), c_0_20]), c_0_27]), c_0_28]), c_0_26])).
fof(c_0_47, plain, ![X33, X34, X35, X36]:(v2_struct_0(X33)|(~v2_pre_topc(X33)|~l1_pre_topc(X33))|(v2_struct_0(X34)|(~v4_orders_2(X34)|(~v7_waybel_0(X34)|~l1_waybel_0(X34,X33)))|(~m2_yellow_6(X35,X33,X34)|(~m1_subset_1(X36,u1_struct_0(X33))|(~r3_waybel_9(X33,X35,X36)|r3_waybel_9(X33,X34,X36)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t31_waybel_9])])])])).
cnf(c_0_48, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r3_waybel_9(X1,X2,X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k10_yellow_6(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_49, plain, (r1_tarski(X1,X2)|m1_subset_1(esk6_2(X1,X2),X3)|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_50, negated_conjecture, (m1_subset_1(k10_yellow_6(esk1_0,esk3_1(esk5_1(esk1_0))),k1_zfmisc_1(u1_struct_0(esk1_0)))|v1_compts_1(esk1_0)|v2_struct_0(esk5_1(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_19]), c_0_18])]), c_0_20]), c_0_44]), c_0_45]), c_0_46])).
cnf(c_0_51, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r3_waybel_9(X1,X2,X4)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~m2_yellow_6(X3,X1,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~r3_waybel_9(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_52, plain, (r1_tarski(k10_yellow_6(X1,X2),X3)|r3_waybel_9(X1,X2,esk6_2(k10_yellow_6(X1,X2),X3))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk6_2(k10_yellow_6(X1,X2),X3),u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_48, c_0_41])).
cnf(c_0_53, negated_conjecture, (r1_tarski(k10_yellow_6(esk1_0,esk3_1(esk5_1(esk1_0))),X1)|m1_subset_1(esk6_2(k10_yellow_6(esk1_0,esk3_1(esk5_1(esk1_0))),X1),u1_struct_0(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(esk5_1(esk1_0))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
fof(c_0_54, plain, ![X37, X38]:(((r1_tarski(X37,X38)|X37!=X38)&(r1_tarski(X38,X37)|X37!=X38))&(~r1_tarski(X37,X38)|~r1_tarski(X38,X37)|X37=X38)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_55, plain, ![X39]:r1_tarski(k1_xboole_0,X39), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_56, negated_conjecture, (r3_waybel_9(esk1_0,esk5_1(esk1_0),X1)|v1_compts_1(esk1_0)|v2_struct_0(esk5_1(esk1_0))|~r3_waybel_9(esk1_0,esk3_1(esk5_1(esk1_0)),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_34]), c_0_19]), c_0_18])]), c_0_20]), c_0_27]), c_0_28]), c_0_26])).
cnf(c_0_57, negated_conjecture, (r1_tarski(k10_yellow_6(esk1_0,esk3_1(esk5_1(esk1_0))),X1)|r3_waybel_9(esk1_0,esk3_1(esk5_1(esk1_0)),esk6_2(k10_yellow_6(esk1_0,esk3_1(esk5_1(esk1_0))),X1))|v1_compts_1(esk1_0)|v2_struct_0(esk5_1(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_19]), c_0_18])]), c_0_20]), c_0_44]), c_0_45]), c_0_43]), c_0_46])).
fof(c_0_58, plain, ![X40, X41]:((~v3_yellow_6(X41,X40)|k10_yellow_6(X40,X41)!=k1_xboole_0|(v2_struct_0(X41)|(~v4_orders_2(X41)|(~v7_waybel_0(X41)|~l1_waybel_0(X41,X40))))|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40))))&(k10_yellow_6(X40,X41)=k1_xboole_0|v3_yellow_6(X41,X40)|(v2_struct_0(X41)|(~v4_orders_2(X41)|(~v7_waybel_0(X41)|~l1_waybel_0(X41,X40))))|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d19_yellow_6])])])])])).
cnf(c_0_59, negated_conjecture, (v3_yellow_6(esk3_1(X1),esk1_0)|v2_struct_0(X1)|v1_compts_1(esk1_0)|~v4_orders_2(X1)|~v7_waybel_0(X1)|~l1_waybel_0(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_60, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_61, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_62, plain, (v1_compts_1(X2)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r3_waybel_9(X2,esk5_1(X2),X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_63, negated_conjecture, (r1_tarski(k10_yellow_6(esk1_0,esk3_1(esk5_1(esk1_0))),X1)|r3_waybel_9(esk1_0,esk5_1(esk1_0),esk6_2(k10_yellow_6(esk1_0,esk3_1(esk5_1(esk1_0))),X1))|v1_compts_1(esk1_0)|v2_struct_0(esk5_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_53])).
cnf(c_0_64, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~v3_yellow_6(X1,X2)|k10_yellow_6(X2,X1)!=k1_xboole_0|~v4_orders_2(X1)|~v7_waybel_0(X1)|~l1_waybel_0(X1,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_65, negated_conjecture, (v3_yellow_6(esk3_1(esk5_1(esk1_0)),esk1_0)|v1_compts_1(esk1_0)|v2_struct_0(esk5_1(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_26]), c_0_27]), c_0_28])).
cnf(c_0_66, plain, (k1_xboole_0=X1|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_67, negated_conjecture, (r1_tarski(k10_yellow_6(esk1_0,esk3_1(esk5_1(esk1_0))),X1)|v1_compts_1(esk1_0)|v2_struct_0(esk5_1(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_19]), c_0_18])]), c_0_20]), c_0_53])).
cnf(c_0_68, negated_conjecture, (v1_compts_1(esk1_0)|v2_struct_0(esk5_1(esk1_0))|k10_yellow_6(esk1_0,esk3_1(esk5_1(esk1_0)))!=k1_xboole_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_19]), c_0_18])]), c_0_20]), c_0_44]), c_0_45]), c_0_43]), c_0_46])).
cnf(c_0_69, plain, (v1_compts_1(X1)|v2_struct_0(X1)|~v2_struct_0(esk5_1(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_70, negated_conjecture, (v1_compts_1(esk1_0)|v2_struct_0(esk5_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68])).
cnf(c_0_71, plain, (m1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v1_compts_1(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_72, negated_conjecture, (v1_compts_1(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_19]), c_0_18])]), c_0_20])).
cnf(c_0_73, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_74, negated_conjecture, (v7_waybel_0(esk2_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_75, negated_conjecture, (v4_orders_2(esk2_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_76, negated_conjecture, (~v2_struct_0(esk2_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_77, plain, ![X42, X43, X44]:((m2_yellow_6(esk7_3(X42,X43,X44),X42,X43)|~r3_waybel_9(X42,X43,X44)|~m1_subset_1(X44,u1_struct_0(X42))|(v2_struct_0(X43)|(~v4_orders_2(X43)|(~v7_waybel_0(X43)|~l1_waybel_0(X43,X42))))|(v2_struct_0(X42)|(~v2_pre_topc(X42)|~l1_pre_topc(X42))))&(r2_hidden(X44,k10_yellow_6(X42,esk7_3(X42,X43,X44)))|~r3_waybel_9(X42,X43,X44)|~m1_subset_1(X44,u1_struct_0(X42))|(v2_struct_0(X43)|(~v4_orders_2(X43)|(~v7_waybel_0(X43)|~l1_waybel_0(X43,X42))))|(v2_struct_0(X42)|(~v2_pre_topc(X42)|~l1_pre_topc(X42))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t32_waybel_9])])])])])])).
cnf(c_0_78, negated_conjecture, (m1_subset_1(esk4_2(esk1_0,X1),u1_struct_0(esk1_0))|v2_struct_0(X1)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_19]), c_0_18])]), c_0_20])).
cnf(c_0_79, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_72])])).
cnf(c_0_80, negated_conjecture, (v7_waybel_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_72])])).
cnf(c_0_81, negated_conjecture, (v4_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_72])])).
cnf(c_0_82, negated_conjecture, (~v2_struct_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_72])])).
cnf(c_0_83, plain, (r3_waybel_9(X1,X2,esk4_2(X1,X2))|v2_struct_0(X2)|v2_struct_0(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v1_compts_1(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_84, plain, (m2_yellow_6(esk7_3(X1,X2,X3),X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_85, negated_conjecture, (m1_subset_1(esk4_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80]), c_0_81])]), c_0_82])).
cnf(c_0_86, negated_conjecture, (r3_waybel_9(esk1_0,X1,esk4_2(esk1_0,X1))|v2_struct_0(X1)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_72]), c_0_19]), c_0_18])]), c_0_20])).
cnf(c_0_87, negated_conjecture, (m2_yellow_6(esk7_3(esk1_0,X1,esk4_2(esk1_0,esk2_0)),esk1_0,X1)|v2_struct_0(X1)|~r3_waybel_9(esk1_0,X1,esk4_2(esk1_0,esk2_0))|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_19]), c_0_18])]), c_0_20])).
cnf(c_0_88, negated_conjecture, (r3_waybel_9(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_79]), c_0_80]), c_0_81])]), c_0_82])).
cnf(c_0_89, plain, (r2_hidden(X1,k10_yellow_6(X2,esk7_3(X2,X3,X1)))|v2_struct_0(X3)|v2_struct_0(X2)|~r3_waybel_9(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v4_orders_2(X3)|~v7_waybel_0(X3)|~l1_waybel_0(X3,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_90, negated_conjecture, (m2_yellow_6(esk7_3(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0)),esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_79]), c_0_80]), c_0_81])]), c_0_82])).
fof(c_0_91, plain, ![X46, X47]:(~r2_hidden(X46,X47)|~r1_tarski(X47,X46)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_92, negated_conjecture, (r2_hidden(esk4_2(esk1_0,esk2_0),k10_yellow_6(esk1_0,esk7_3(esk1_0,X1,esk4_2(esk1_0,esk2_0))))|v2_struct_0(X1)|~r3_waybel_9(esk1_0,X1,esk4_2(esk1_0,esk2_0))|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_85]), c_0_19]), c_0_18])]), c_0_20])).
cnf(c_0_93, negated_conjecture, (~m2_yellow_6(X1,esk1_0,esk2_0)|~v3_yellow_6(X1,esk1_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_94, plain, (k10_yellow_6(X1,X2)=k1_xboole_0|v3_yellow_6(X2,X1)|v2_struct_0(X2)|v2_struct_0(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_95, negated_conjecture, (l1_waybel_0(esk7_3(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0)),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_90]), c_0_35]), c_0_79]), c_0_80]), c_0_81])]), c_0_20]), c_0_82])).
cnf(c_0_96, negated_conjecture, (v7_waybel_0(esk7_3(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_90]), c_0_35]), c_0_79]), c_0_80]), c_0_81])]), c_0_20]), c_0_82])).
cnf(c_0_97, negated_conjecture, (v4_orders_2(esk7_3(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_90]), c_0_35]), c_0_79]), c_0_80]), c_0_81])]), c_0_20]), c_0_82])).
cnf(c_0_98, negated_conjecture, (~v2_struct_0(esk7_3(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_90]), c_0_35]), c_0_79]), c_0_80]), c_0_81])]), c_0_20]), c_0_82])).
cnf(c_0_99, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_100, negated_conjecture, (r2_hidden(esk4_2(esk1_0,esk2_0),k10_yellow_6(esk1_0,esk7_3(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_88]), c_0_79]), c_0_80]), c_0_81])]), c_0_82])).
cnf(c_0_101, negated_conjecture, (~v3_yellow_6(X1,esk1_0)|~m2_yellow_6(X1,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_72])])).
cnf(c_0_102, negated_conjecture, (k10_yellow_6(esk1_0,esk7_3(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0)))=k1_xboole_0|v3_yellow_6(esk7_3(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0)),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_96]), c_0_97]), c_0_19]), c_0_18])]), c_0_20]), c_0_98])).
cnf(c_0_103, negated_conjecture, (~r1_tarski(k10_yellow_6(esk1_0,esk7_3(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0))),esk4_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_104, negated_conjecture, (k10_yellow_6(esk1_0,esk7_3(esk1_0,esk2_0,esk4_2(esk1_0,esk2_0)))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_90])])).
cnf(c_0_105, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_104]), c_0_61])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 106
# Proof object clause steps            : 77
# Proof object formula steps           : 29
# Proof object conjectures             : 53
# Proof object clause conjectures      : 50
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 34
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 37
# Proof object simplifying inferences  : 162
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 38
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 38
# Processed clauses                    : 214
# ...of these trivial                  : 6
# ...subsumed                          : 51
# ...remaining for further processing  : 157
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 70
# Backward-rewritten                   : 14
# Generated clauses                    : 275
# ...of the previous two non-trivial   : 247
# Contextual simplify-reflections      : 174
# Paramodulations                      : 273
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 71
#    Positive orientable unit clauses  : 19
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 49
# Current number of unprocessed clauses: 19
# ...number of literals in the above   : 67
# Current number of archived formulas  : 0
# Current number of archived clauses   : 84
# Clause-clause subsumption calls (NU) : 9745
# Rec. Clause-clause subsumption calls : 1214
# Non-unit clause-clause subsumptions  : 294
# Unit Clause-clause subsumption calls : 146
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 14224

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.003 s
# Total time               : 0.019 s
# Maximum resident set size: 28684 pages
