# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t165_funct_1, conjecture, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((r1_tarski(X3,k1_relat_1(X1))&r1_tarski(X3,k1_relat_1(X2)))=>(k7_relat_1(X1,X3)=k7_relat_1(X2,X3)<=>![X4]:(r2_hidden(X4,X3)=>k1_funct_1(X1,X4)=k1_funct_1(X2,X4)))))), file('data/MPTP/proofs/52fe3d5e76804ec0af69f65147c07383.E_input', t165_funct_1)).
fof(t68_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X2)=k3_xboole_0(k1_relat_1(X3),X1)&![X4]:(r2_hidden(X4,k1_relat_1(X2))=>k1_funct_1(X2,X4)=k1_funct_1(X3,X4)))=>X2=k7_relat_1(X3,X1)))), file('data/MPTP/proofs/52fe3d5e76804ec0af69f65147c07383.E_input', t68_funct_1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/52fe3d5e76804ec0af69f65147c07383.E_input', t12_setfam_1)).
fof(t90_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k1_relat_1(k7_relat_1(X2,X1))=k3_xboole_0(k1_relat_1(X2),X1)), file('data/MPTP/proofs/52fe3d5e76804ec0af69f65147c07383.E_input', t90_relat_1)).
fof(t91_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(X1,k1_relat_1(X2))=>k1_relat_1(k7_relat_1(X2,X1))=X1)), file('data/MPTP/proofs/52fe3d5e76804ec0af69f65147c07383.E_input', t91_relat_1)).
fof(fc8_funct_1, axiom, ![X1, X2]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_relat_1(k7_relat_1(X1,X2))&v1_funct_1(k7_relat_1(X1,X2)))), file('data/MPTP/proofs/52fe3d5e76804ec0af69f65147c07383.E_input', fc8_funct_1)).
fof(t72_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,X2)=>k1_funct_1(k7_relat_1(X3,X2),X1)=k1_funct_1(X3,X1))), file('data/MPTP/proofs/52fe3d5e76804ec0af69f65147c07383.E_input', t72_funct_1)).
fof(c_0_7, negated_conjecture, ~(![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((r1_tarski(X3,k1_relat_1(X1))&r1_tarski(X3,k1_relat_1(X2)))=>(k7_relat_1(X1,X3)=k7_relat_1(X2,X3)<=>![X4]:(r2_hidden(X4,X3)=>k1_funct_1(X1,X4)=k1_funct_1(X2,X4))))))), inference(assume_negation,[status(cth)],[t165_funct_1])).
fof(c_0_8, plain, ![X10, X11, X12]:((r2_hidden(esk5_3(X10,X11,X12),k1_relat_1(X11))|k1_relat_1(X11)!=k3_xboole_0(k1_relat_1(X12),X10)|X11=k7_relat_1(X12,X10)|(~v1_relat_1(X12)|~v1_funct_1(X12))|(~v1_relat_1(X11)|~v1_funct_1(X11)))&(k1_funct_1(X11,esk5_3(X10,X11,X12))!=k1_funct_1(X12,esk5_3(X10,X11,X12))|k1_relat_1(X11)!=k3_xboole_0(k1_relat_1(X12),X10)|X11=k7_relat_1(X12,X10)|(~v1_relat_1(X12)|~v1_funct_1(X12))|(~v1_relat_1(X11)|~v1_funct_1(X11)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t68_funct_1])])])])])).
fof(c_0_9, plain, ![X14, X15]:k1_setfam_1(k2_tarski(X14,X15))=k3_xboole_0(X14,X15), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_10, plain, ![X16, X17]:(~v1_relat_1(X17)|k1_relat_1(k7_relat_1(X17,X16))=k3_xboole_0(k1_relat_1(X17),X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t90_relat_1])])).
fof(c_0_11, plain, ![X18, X19]:(~v1_relat_1(X19)|(~r1_tarski(X18,k1_relat_1(X19))|k1_relat_1(k7_relat_1(X19,X18))=X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t91_relat_1])])).
fof(c_0_12, negated_conjecture, ![X9]:((v1_relat_1(esk1_0)&v1_funct_1(esk1_0))&((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&((r1_tarski(esk3_0,k1_relat_1(esk1_0))&r1_tarski(esk3_0,k1_relat_1(esk2_0)))&(((r2_hidden(esk4_0,esk3_0)|k7_relat_1(esk1_0,esk3_0)!=k7_relat_1(esk2_0,esk3_0))&(k1_funct_1(esk1_0,esk4_0)!=k1_funct_1(esk2_0,esk4_0)|k7_relat_1(esk1_0,esk3_0)!=k7_relat_1(esk2_0,esk3_0)))&(k7_relat_1(esk1_0,esk3_0)=k7_relat_1(esk2_0,esk3_0)|(~r2_hidden(X9,esk3_0)|k1_funct_1(esk1_0,X9)=k1_funct_1(esk2_0,X9))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).
fof(c_0_13, plain, ![X20, X21]:((v1_relat_1(k7_relat_1(X20,X21))|(~v1_relat_1(X20)|~v1_funct_1(X20)))&(v1_funct_1(k7_relat_1(X20,X21))|(~v1_relat_1(X20)|~v1_funct_1(X20)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc8_funct_1])])])).
cnf(c_0_14, plain, (X1=k7_relat_1(X3,X2)|k1_funct_1(X1,esk5_3(X2,X1,X3))!=k1_funct_1(X3,esk5_3(X2,X1,X3))|k1_relat_1(X1)!=k3_xboole_0(k1_relat_1(X3),X2)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, plain, (k1_relat_1(k7_relat_1(X1,X2))=k3_xboole_0(k1_relat_1(X1),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, plain, (r2_hidden(esk5_3(X1,X2,X3),k1_relat_1(X2))|X2=k7_relat_1(X3,X1)|k1_relat_1(X2)!=k3_xboole_0(k1_relat_1(X3),X1)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_18, plain, (k1_relat_1(k7_relat_1(X1,X2))=X2|~v1_relat_1(X1)|~r1_tarski(X2,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19, negated_conjecture, (r1_tarski(esk3_0,k1_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, plain, (v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23, plain, (X1=k7_relat_1(X3,X2)|k1_relat_1(X1)!=k1_setfam_1(k2_tarski(k1_relat_1(X3),X2))|k1_funct_1(X1,esk5_3(X2,X1,X3))!=k1_funct_1(X3,esk5_3(X2,X1,X3))|~v1_relat_1(X3)|~v1_relat_1(X1)|~v1_funct_1(X3)|~v1_funct_1(X1)), inference(rw,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_24, plain, (k1_relat_1(k7_relat_1(X1,X2))=k1_setfam_1(k2_tarski(k1_relat_1(X1),X2))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_16, c_0_15])).
fof(c_0_25, plain, ![X22, X23, X24]:(~v1_relat_1(X24)|~v1_funct_1(X24)|(~r2_hidden(X22,X23)|k1_funct_1(k7_relat_1(X24,X23),X22)=k1_funct_1(X24,X22))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t72_funct_1])])).
cnf(c_0_26, plain, (X2=k7_relat_1(X3,X1)|r2_hidden(esk5_3(X1,X2,X3),k1_relat_1(X2))|k1_relat_1(X2)!=k1_setfam_1(k2_tarski(k1_relat_1(X3),X1))|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2)), inference(rw,[status(thm)],[c_0_17, c_0_15])).
cnf(c_0_27, negated_conjecture, (k1_relat_1(k7_relat_1(esk2_0,esk3_0))=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])])).
cnf(c_0_28, negated_conjecture, (v1_relat_1(k7_relat_1(esk2_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_20])])).
cnf(c_0_29, plain, (X1=k7_relat_1(X2,X3)|k1_funct_1(X1,esk5_3(X3,X1,X2))!=k1_funct_1(X2,esk5_3(X3,X1,X2))|k1_relat_1(X1)!=k1_relat_1(k7_relat_1(X2,X3))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_30, plain, (k1_funct_1(k7_relat_1(X1,X3),X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, plain, (v1_funct_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_32, negated_conjecture, (k7_relat_1(esk2_0,esk3_0)=k7_relat_1(X1,X2)|r2_hidden(esk5_3(X2,k7_relat_1(esk2_0,esk3_0),X1),esk3_0)|k1_setfam_1(k2_tarski(k1_relat_1(X1),X2))!=esk3_0|~v1_funct_1(k7_relat_1(esk2_0,esk3_0))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])])).
cnf(c_0_33, plain, (k7_relat_1(X1,X2)=k7_relat_1(X3,X4)|k1_funct_1(X1,esk5_3(X4,k7_relat_1(X1,X2),X3))!=k1_funct_1(X3,esk5_3(X4,k7_relat_1(X1,X2),X3))|k1_relat_1(k7_relat_1(X1,X2))!=k1_relat_1(k7_relat_1(X3,X4))|~r2_hidden(esk5_3(X4,k7_relat_1(X1,X2),X3),X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_21]), c_0_31])).
cnf(c_0_34, negated_conjecture, (k7_relat_1(esk2_0,esk3_0)=k7_relat_1(X1,X2)|r2_hidden(esk5_3(X2,k7_relat_1(esk2_0,esk3_0),X1),esk3_0)|k1_relat_1(k7_relat_1(X1,X2))!=esk3_0|~v1_funct_1(k7_relat_1(esk2_0,esk3_0))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_32, c_0_24])).
cnf(c_0_35, negated_conjecture, (k7_relat_1(esk1_0,esk3_0)=k7_relat_1(esk2_0,esk3_0)|k1_funct_1(esk1_0,X1)=k1_funct_1(esk2_0,X1)|~r2_hidden(X1,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_36, negated_conjecture, (k7_relat_1(esk2_0,esk3_0)=k7_relat_1(X1,X2)|k1_funct_1(esk2_0,esk5_3(X2,k7_relat_1(esk2_0,esk3_0),X1))!=k1_funct_1(X1,esk5_3(X2,k7_relat_1(esk2_0,esk3_0),X1))|k1_relat_1(k7_relat_1(X1,X2))!=esk3_0|~v1_funct_1(k7_relat_1(esk2_0,esk3_0))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_27]), c_0_22]), c_0_20])])).
cnf(c_0_37, negated_conjecture, (k1_funct_1(esk1_0,esk5_3(X1,k7_relat_1(esk2_0,esk3_0),X2))=k1_funct_1(esk2_0,esk5_3(X1,k7_relat_1(esk2_0,esk3_0),X2))|k7_relat_1(esk1_0,esk3_0)=k7_relat_1(esk2_0,esk3_0)|k7_relat_1(esk2_0,esk3_0)=k7_relat_1(X2,X1)|k1_relat_1(k7_relat_1(X2,X1))!=esk3_0|~v1_funct_1(k7_relat_1(esk2_0,esk3_0))|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_35, c_0_34])).
cnf(c_0_38, negated_conjecture, (v1_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_39, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_40, negated_conjecture, (k7_relat_1(esk1_0,esk3_0)=k7_relat_1(esk2_0,esk3_0)|k7_relat_1(esk1_0,X1)=k7_relat_1(esk2_0,esk3_0)|k1_relat_1(k7_relat_1(esk1_0,X1))!=esk3_0|~v1_funct_1(k7_relat_1(esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_39])])).
cnf(c_0_41, negated_conjecture, (r1_tarski(esk3_0,k1_relat_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_42, negated_conjecture, (k7_relat_1(esk1_0,esk3_0)=k7_relat_1(esk2_0,esk3_0)|k7_relat_1(esk1_0,X1)=k7_relat_1(esk2_0,esk3_0)|k1_relat_1(k7_relat_1(esk1_0,X1))!=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_31]), c_0_22]), c_0_20])])).
cnf(c_0_43, negated_conjecture, (k1_relat_1(k7_relat_1(esk1_0,esk3_0))=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_41]), c_0_39])])).
cnf(c_0_44, negated_conjecture, (k7_relat_1(esk1_0,esk3_0)=k7_relat_1(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_45, negated_conjecture, (k1_funct_1(k7_relat_1(esk2_0,esk3_0),X1)=k1_funct_1(esk1_0,X1)|~r2_hidden(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_44]), c_0_38]), c_0_39])])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk4_0,esk3_0)|k7_relat_1(esk1_0,esk3_0)!=k7_relat_1(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_47, negated_conjecture, (k1_funct_1(esk1_0,esk4_0)!=k1_funct_1(esk2_0,esk4_0)|k7_relat_1(esk1_0,esk3_0)!=k7_relat_1(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_48, negated_conjecture, (k1_funct_1(esk1_0,X1)=k1_funct_1(esk2_0,X1)|~r2_hidden(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_45]), c_0_22]), c_0_20])])).
cnf(c_0_49, negated_conjecture, (r2_hidden(esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_44])])).
cnf(c_0_50, negated_conjecture, (k1_funct_1(esk1_0,esk4_0)!=k1_funct_1(esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_44])])).
cnf(c_0_51, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 52
# Proof object clause steps            : 37
# Proof object formula steps           : 15
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 34
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 16
# Processed clauses                    : 129
# ...of these trivial                  : 1
# ...subsumed                          : 21
# ...remaining for further processing  : 107
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 47
# Generated clauses                    : 142
# ...of the previous two non-trivial   : 178
# Contextual simplify-reflections      : 12
# Paramodulations                      : 132
# Factorizations                       : 0
# Equation resolutions                 : 10
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 38
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 23
# Current number of unprocessed clauses: 81
# ...number of literals in the above   : 537
# Current number of archived formulas  : 0
# Current number of archived clauses   : 70
# Clause-clause subsumption calls (NU) : 1189
# Rec. Clause-clause subsumption calls : 172
# Non-unit clause-clause subsumptions  : 39
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8736

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.003 s
# Total time               : 0.014 s
# Maximum resident set size: 28712 pages
