# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/538957f51947420695afadeb2be33b08.E_input', d5_xboole_0)).
fof(t34_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4, X5, X6]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k4_mcart_1(X3,X4,X5,X6))))))), file('data/MPTP/proofs/538957f51947420695afadeb2be33b08.E_input', t34_mcart_1)).
fof(t65_funct_2, conjecture, ![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,k1_tarski(X2))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_tarski(X2))))))=>(r2_hidden(X3,X1)=>k1_funct_1(X4,X3)=X2)), file('data/MPTP/proofs/538957f51947420695afadeb2be33b08.E_input', t65_funct_2)).
fof(l33_zfmisc_1, axiom, ![X1, X2]:(k4_xboole_0(k1_tarski(X1),X2)=k1_tarski(X1)<=>~(r2_hidden(X1,X2))), file('data/MPTP/proofs/538957f51947420695afadeb2be33b08.E_input', l33_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/538957f51947420695afadeb2be33b08.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/538957f51947420695afadeb2be33b08.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/538957f51947420695afadeb2be33b08.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/538957f51947420695afadeb2be33b08.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/538957f51947420695afadeb2be33b08.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/538957f51947420695afadeb2be33b08.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/538957f51947420695afadeb2be33b08.E_input', t75_enumset1)).
fof(d6_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:(X9=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)<=>![X10]:(r2_hidden(X10,X9)<=>~((~(X10=X1)&(~(X10=X2)&(~(X10=X3)&(~(X10=X4)&(~(X10=X5)&(~(X10=X6)&(~(X10=X7)&~(X10=X8))))))))))), file('data/MPTP/proofs/538957f51947420695afadeb2be33b08.E_input', d6_enumset1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/538957f51947420695afadeb2be33b08.E_input', d1_tarski)).
fof(t7_funct_2, axiom, ![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(r2_hidden(X3,X1)=>(X2=k1_xboole_0|r2_hidden(k1_funct_1(X4,X3),X2)))), file('data/MPTP/proofs/538957f51947420695afadeb2be33b08.E_input', t7_funct_2)).
fof(c_0_14, plain, ![X79, X80, X81, X82, X83, X84, X85, X86]:((((r2_hidden(X82,X79)|~r2_hidden(X82,X81)|X81!=k4_xboole_0(X79,X80))&(~r2_hidden(X82,X80)|~r2_hidden(X82,X81)|X81!=k4_xboole_0(X79,X80)))&(~r2_hidden(X83,X79)|r2_hidden(X83,X80)|r2_hidden(X83,X81)|X81!=k4_xboole_0(X79,X80)))&((~r2_hidden(esk8_3(X84,X85,X86),X86)|(~r2_hidden(esk8_3(X84,X85,X86),X84)|r2_hidden(esk8_3(X84,X85,X86),X85))|X86=k4_xboole_0(X84,X85))&((r2_hidden(esk8_3(X84,X85,X86),X84)|r2_hidden(esk8_3(X84,X85,X86),X86)|X86=k4_xboole_0(X84,X85))&(~r2_hidden(esk8_3(X84,X85,X86),X85)|r2_hidden(esk8_3(X84,X85,X86),X86)|X86=k4_xboole_0(X84,X85))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_15, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_16, plain, ![X15, X17, X18, X19, X20]:((r2_hidden(esk5_1(X15),X15)|X15=k1_xboole_0)&((~r2_hidden(X17,X15)|esk5_1(X15)!=k4_mcart_1(X17,X18,X19,X20)|X15=k1_xboole_0)&(~r2_hidden(X18,X15)|esk5_1(X15)!=k4_mcart_1(X17,X18,X19,X20)|X15=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_mcart_1])])])])])])])).
cnf(c_0_17, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, negated_conjecture, ~(![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,k1_tarski(X2))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_tarski(X2))))))=>(r2_hidden(X3,X1)=>k1_funct_1(X4,X3)=X2))), inference(assume_negation,[status(cth)],[t65_funct_2])).
fof(c_0_19, plain, ![X49, X50]:((k4_xboole_0(k1_tarski(X49),X50)!=k1_tarski(X49)|~r2_hidden(X49,X50))&(r2_hidden(X49,X50)|k4_xboole_0(k1_tarski(X49),X50)=k1_tarski(X49))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l33_zfmisc_1])])])).
fof(c_0_20, plain, ![X67]:k2_tarski(X67,X67)=k1_tarski(X67), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_21, plain, ![X51, X52]:k1_enumset1(X51,X51,X52)=k2_tarski(X51,X52), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_22, plain, ![X46, X47, X48]:k2_enumset1(X46,X46,X47,X48)=k1_enumset1(X46,X47,X48), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_23, plain, ![X42, X43, X44, X45]:k3_enumset1(X42,X42,X43,X44,X45)=k2_enumset1(X42,X43,X44,X45), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_24, plain, ![X74, X75, X76, X77, X78]:k4_enumset1(X74,X74,X75,X76,X77,X78)=k3_enumset1(X74,X75,X76,X77,X78), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_25, plain, ![X68, X69, X70, X71, X72, X73]:k5_enumset1(X68,X68,X69,X70,X71,X72,X73)=k4_enumset1(X68,X69,X70,X71,X72,X73), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_26, plain, ![X53, X54, X55, X56, X57, X58, X59]:k6_enumset1(X53,X53,X54,X55,X56,X57,X58,X59)=k5_enumset1(X53,X54,X55,X56,X57,X58,X59), inference(variable_rename,[status(thm)],[t75_enumset1])).
cnf(c_0_27, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_15])).
cnf(c_0_28, plain, (r2_hidden(esk5_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_29, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_17])).
fof(c_0_30, plain, ![X21, X22, X23, X24, X25, X26, X27, X28, X29, X30, X31, X32, X33, X34, X35, X36, X37, X38, X39, X40]:(((~r2_hidden(X30,X29)|(X30=X21|(X30=X22|(X30=X23|(X30=X24|(X30=X25|(X30=X26|(X30=X27|X30=X28)))))))|X29!=k6_enumset1(X21,X22,X23,X24,X25,X26,X27,X28))&((X31!=X21|r2_hidden(X31,X29)|X29!=k6_enumset1(X21,X22,X23,X24,X25,X26,X27,X28))&((X31!=X22|r2_hidden(X31,X29)|X29!=k6_enumset1(X21,X22,X23,X24,X25,X26,X27,X28))&((X31!=X23|r2_hidden(X31,X29)|X29!=k6_enumset1(X21,X22,X23,X24,X25,X26,X27,X28))&((X31!=X24|r2_hidden(X31,X29)|X29!=k6_enumset1(X21,X22,X23,X24,X25,X26,X27,X28))&((X31!=X25|r2_hidden(X31,X29)|X29!=k6_enumset1(X21,X22,X23,X24,X25,X26,X27,X28))&((X31!=X26|r2_hidden(X31,X29)|X29!=k6_enumset1(X21,X22,X23,X24,X25,X26,X27,X28))&((X31!=X27|r2_hidden(X31,X29)|X29!=k6_enumset1(X21,X22,X23,X24,X25,X26,X27,X28))&(X31!=X28|r2_hidden(X31,X29)|X29!=k6_enumset1(X21,X22,X23,X24,X25,X26,X27,X28))))))))))&(((esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40)!=X32|~r2_hidden(esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40),X40)|X40=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40)!=X33|~r2_hidden(esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40),X40)|X40=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40)!=X34|~r2_hidden(esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40),X40)|X40=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40)!=X35|~r2_hidden(esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40),X40)|X40=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40)!=X36|~r2_hidden(esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40),X40)|X40=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40)!=X37|~r2_hidden(esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40),X40)|X40=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40)!=X38|~r2_hidden(esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40),X40)|X40=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&(esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40)!=X39|~r2_hidden(esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40),X40)|X40=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39)))))))))&(r2_hidden(esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40),X40)|(esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40)=X32|(esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40)=X33|(esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40)=X34|(esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40)=X35|(esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40)=X36|(esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40)=X37|(esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40)=X38|esk6_9(X32,X33,X34,X35,X36,X37,X38,X39,X40)=X39)))))))|X40=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_enumset1])])])])])])])).
fof(c_0_31, plain, ![X60, X61, X62, X63, X64, X65]:(((~r2_hidden(X62,X61)|X62=X60|X61!=k1_tarski(X60))&(X63!=X60|r2_hidden(X63,X61)|X61!=k1_tarski(X60)))&((~r2_hidden(esk7_2(X64,X65),X65)|esk7_2(X64,X65)!=X64|X65=k1_tarski(X64))&(r2_hidden(esk7_2(X64,X65),X65)|esk7_2(X64,X65)=X64|X65=k1_tarski(X64)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_32, negated_conjecture, ((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,esk1_0,k1_tarski(esk2_0))&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k1_tarski(esk2_0))))))&(r2_hidden(esk3_0,esk1_0)&k1_funct_1(esk4_0,esk3_0)!=esk2_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).
cnf(c_0_33, plain, (k4_xboole_0(k1_tarski(X1),X2)!=k1_tarski(X1)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_34, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_35, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_36, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_37, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_38, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_39, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_40, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_41, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r2_hidden(esk5_1(k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_42, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|r2_hidden(esk5_1(k4_xboole_0(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_29, c_0_28])).
cnf(c_0_43, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k6_enumset1(X4,X5,X6,X7,X2,X8,X9,X10)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_44, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_45, plain, ![X88, X89, X90, X91]:(~v1_funct_1(X91)|(~v1_funct_2(X91,X88,X89)|~m1_subset_1(X91,k1_zfmisc_1(k2_zfmisc_1(X88,X89))))|(~r2_hidden(X90,X88)|(X89=k1_xboole_0|r2_hidden(k1_funct_1(X91,X90),X89)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_funct_2])])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k1_tarski(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_47, negated_conjecture, (v1_funct_2(esk4_0,esk1_0,k1_tarski(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_48, plain, (k4_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),X2)!=k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_34]), c_0_34]), c_0_35]), c_0_35]), c_0_36]), c_0_36]), c_0_37]), c_0_37]), c_0_38]), c_0_38]), c_0_39]), c_0_39]), c_0_40]), c_0_40])).
cnf(c_0_49, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_50, plain, (r2_hidden(X1,k6_enumset1(X2,X3,X4,X5,X1,X6,X7,X8))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_43])])).
cnf(c_0_51, plain, (X1=X3|X2!=k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40])).
cnf(c_0_52, plain, (X3=k1_xboole_0|r2_hidden(k1_funct_1(X1,X4),X3)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_53, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40])).
cnf(c_0_54, negated_conjecture, (v1_funct_2(esk4_0,esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40])).
cnf(c_0_55, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_56, plain, (k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])])).
cnf(c_0_57, plain, (X1=X2|~r2_hidden(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2))), inference(er,[status(thm)],[c_0_51])).
cnf(c_0_58, negated_conjecture, (r2_hidden(k1_funct_1(esk4_0,X1),k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))|~r2_hidden(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54]), c_0_55])]), c_0_56])).
cnf(c_0_59, negated_conjecture, (k1_funct_1(esk4_0,esk3_0)!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_60, negated_conjecture, (k1_funct_1(esk4_0,X1)=esk2_0|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_61, negated_conjecture, (r2_hidden(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_62, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 63
# Proof object clause steps            : 34
# Proof object formula steps           : 29
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 48
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 46
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 39
# Processed clauses                    : 107
# ...of these trivial                  : 1
# ...subsumed                          : 10
# ...remaining for further processing  : 96
# Other redundant clauses eliminated   : 23
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 106
# ...of the previous two non-trivial   : 74
# Contextual simplify-reflections      : 1
# Paramodulations                      : 90
# Factorizations                       : 2
# Equation resolutions                 : 23
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 44
#    Positive orientable unit clauses  : 15
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 26
# Current number of unprocessed clauses: 44
# ...number of literals in the above   : 127
# Current number of archived formulas  : 0
# Current number of archived clauses   : 45
# Clause-clause subsumption calls (NU) : 218
# Rec. Clause-clause subsumption calls : 151
# Non-unit clause-clause subsumptions  : 8
# Unit Clause-clause subsumption calls : 6
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 74
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4010

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28716 pages
