# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/543f9fae819f403d9293da55ceca4c27.E_input', d10_xboole_0)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/543f9fae819f403d9293da55ceca4c27.E_input', d5_xboole_0)).
fof(l1_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/543f9fae819f403d9293da55ceca4c27.E_input', l1_zfmisc_1)).
fof(t106_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k4_xboole_0(X2,X3))=>(r1_tarski(X1,X2)&r1_xboole_0(X1,X3))), file('data/MPTP/proofs/543f9fae819f403d9293da55ceca4c27.E_input', t106_xboole_1)).
fof(t65_zfmisc_1, axiom, ![X1, X2]:(k4_xboole_0(X1,k1_tarski(X2))=X1<=>~(r2_hidden(X2,X1))), file('data/MPTP/proofs/543f9fae819f403d9293da55ceca4c27.E_input', t65_zfmisc_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/543f9fae819f403d9293da55ceca4c27.E_input', d1_xboole_0)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/543f9fae819f403d9293da55ceca4c27.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/543f9fae819f403d9293da55ceca4c27.E_input', t70_enumset1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/543f9fae819f403d9293da55ceca4c27.E_input', t2_xboole_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/543f9fae819f403d9293da55ceca4c27.E_input', t7_xboole_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/543f9fae819f403d9293da55ceca4c27.E_input', commutativity_k2_tarski)).
fof(t103_zfmisc_1, conjecture, ![X1, X2, X3, X4]:~((r1_tarski(X1,k2_zfmisc_1(X2,X3))&(r2_hidden(X4,X1)&![X5, X6]:~((r2_hidden(X5,X2)&(r2_hidden(X6,X3)&X4=k4_tarski(X5,X6))))))), file('data/MPTP/proofs/543f9fae819f403d9293da55ceca4c27.E_input', t103_zfmisc_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/543f9fae819f403d9293da55ceca4c27.E_input', t3_boole)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/543f9fae819f403d9293da55ceca4c27.E_input', d4_xboole_0)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/543f9fae819f403d9293da55ceca4c27.E_input', t48_xboole_1)).
fof(t73_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,k2_xboole_0(X2,X3))&r1_xboole_0(X1,X3))=>r1_tarski(X1,X2)), file('data/MPTP/proofs/543f9fae819f403d9293da55ceca4c27.E_input', t73_xboole_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/543f9fae819f403d9293da55ceca4c27.E_input', t7_xboole_0)).
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/543f9fae819f403d9293da55ceca4c27.E_input', d2_zfmisc_1)).
fof(t63_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_xboole_0(X2,X3))=>r1_xboole_0(X1,X3)), file('data/MPTP/proofs/543f9fae819f403d9293da55ceca4c27.E_input', t63_xboole_1)).
fof(c_0_19, plain, ![X7, X8]:(((r1_tarski(X7,X8)|X7!=X8)&(r1_tarski(X8,X7)|X7!=X8))&(~r1_tarski(X7,X8)|~r1_tarski(X8,X7)|X7=X8)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_20, plain, ![X16, X17, X18, X19, X20, X21, X22, X23]:((((r2_hidden(X19,X16)|~r2_hidden(X19,X18)|X18!=k4_xboole_0(X16,X17))&(~r2_hidden(X19,X17)|~r2_hidden(X19,X18)|X18!=k4_xboole_0(X16,X17)))&(~r2_hidden(X20,X16)|r2_hidden(X20,X17)|r2_hidden(X20,X18)|X18!=k4_xboole_0(X16,X17)))&((~r2_hidden(esk1_3(X21,X22,X23),X23)|(~r2_hidden(esk1_3(X21,X22,X23),X21)|r2_hidden(esk1_3(X21,X22,X23),X22))|X23=k4_xboole_0(X21,X22))&((r2_hidden(esk1_3(X21,X22,X23),X21)|r2_hidden(esk1_3(X21,X22,X23),X23)|X23=k4_xboole_0(X21,X22))&(~r2_hidden(esk1_3(X21,X22,X23),X22)|r2_hidden(esk1_3(X21,X22,X23),X23)|X23=k4_xboole_0(X21,X22))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_21, plain, ![X25, X26]:((~r1_tarski(k1_tarski(X25),X26)|r2_hidden(X25,X26))&(~r2_hidden(X25,X26)|r1_tarski(k1_tarski(X25),X26))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])).
cnf(c_0_22, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_23, plain, ![X13, X14, X15]:((r1_tarski(X13,X14)|~r1_tarski(X13,k4_xboole_0(X14,X15)))&(r1_xboole_0(X13,X15)|~r1_tarski(X13,k4_xboole_0(X14,X15)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t106_xboole_1])])])).
cnf(c_0_24, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_25, plain, ![X47, X48]:((k4_xboole_0(X47,k1_tarski(X48))!=X47|~r2_hidden(X48,X47))&(r2_hidden(X48,X47)|k4_xboole_0(X47,k1_tarski(X48))=X47)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t65_zfmisc_1])])])).
fof(c_0_26, plain, ![X31, X32, X33]:((~v1_xboole_0(X31)|~r2_hidden(X32,X31))&(r2_hidden(esk2_1(X33),X33)|v1_xboole_0(X33))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_27, plain, (r2_hidden(X1,X2)|~r1_tarski(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_22])).
fof(c_0_29, plain, ![X9, X10]:k3_tarski(k2_tarski(X9,X10))=k2_xboole_0(X9,X10), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_30, plain, ![X11, X12]:k1_enumset1(X11,X11,X12)=k2_tarski(X11,X12), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_31, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k4_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_33, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_24])).
cnf(c_0_34, plain, (r2_hidden(X1,X2)|k4_xboole_0(X2,k1_tarski(X1))=X2), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (r2_hidden(X1,k1_tarski(X1))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
fof(c_0_37, plain, ![X27]:r1_tarski(k1_xboole_0,X27), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_38, plain, ![X51, X52]:r1_tarski(X51,k2_xboole_0(X51,X52)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_39, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_40, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_41, plain, ![X76, X77]:k2_tarski(X76,X77)=k2_tarski(X77,X76), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_42, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k4_xboole_0(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_43, negated_conjecture, ~(![X1, X2, X3, X4]:~((r1_tarski(X1,k2_zfmisc_1(X2,X3))&(r2_hidden(X4,X1)&![X5, X6]:~((r2_hidden(X5,X2)&(r2_hidden(X6,X3)&X4=k4_tarski(X5,X6)))))))), inference(assume_negation,[status(cth)],[t103_zfmisc_1])).
cnf(c_0_44, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_45, plain, (r2_hidden(esk2_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_46, plain, ![X46]:k4_xboole_0(X46,k1_xboole_0)=X46, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_47, plain, (r2_hidden(X1,X2)|~r2_hidden(X3,k1_tarski(X1))|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_48, plain, (~v1_xboole_0(k1_tarski(X1))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
fof(c_0_49, plain, ![X35, X36, X37, X38, X39, X40, X41, X42]:((((r2_hidden(X38,X35)|~r2_hidden(X38,X37)|X37!=k3_xboole_0(X35,X36))&(r2_hidden(X38,X36)|~r2_hidden(X38,X37)|X37!=k3_xboole_0(X35,X36)))&(~r2_hidden(X39,X35)|~r2_hidden(X39,X36)|r2_hidden(X39,X37)|X37!=k3_xboole_0(X35,X36)))&((~r2_hidden(esk3_3(X40,X41,X42),X42)|(~r2_hidden(esk3_3(X40,X41,X42),X40)|~r2_hidden(esk3_3(X40,X41,X42),X41))|X42=k3_xboole_0(X40,X41))&((r2_hidden(esk3_3(X40,X41,X42),X40)|r2_hidden(esk3_3(X40,X41,X42),X42)|X42=k3_xboole_0(X40,X41))&(r2_hidden(esk3_3(X40,X41,X42),X41)|r2_hidden(esk3_3(X40,X41,X42),X42)|X42=k3_xboole_0(X40,X41))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_50, plain, ![X44, X45]:k4_xboole_0(X44,k4_xboole_0(X44,X45))=k3_xboole_0(X44,X45), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_51, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_52, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_53, plain, ![X28, X29, X30]:(~r1_tarski(X28,k2_xboole_0(X29,X30))|~r1_xboole_0(X28,X30)|r1_tarski(X28,X29)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t73_xboole_1])])).
cnf(c_0_54, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_55, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_56, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_57, plain, (r2_hidden(X1,X2)|r1_xboole_0(X3,k1_tarski(X1))|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_42, c_0_34])).
fof(c_0_58, negated_conjecture, ![X57, X58]:(r1_tarski(esk5_0,k2_zfmisc_1(esk6_0,esk7_0))&(r2_hidden(esk8_0,esk5_0)&(~r2_hidden(X57,esk6_0)|(~r2_hidden(X58,esk7_0)|esk8_0!=k4_tarski(X57,X58))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_43])])])])).
cnf(c_0_59, plain, (v1_xboole_0(k4_xboole_0(X1,X2))|r1_tarski(k1_tarski(esk2_1(k4_xboole_0(X1,X2))),X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_60, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_61, plain, (r2_hidden(X1,X2)|~r2_hidden(esk2_1(k1_tarski(X1)),X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_45]), c_0_48])).
cnf(c_0_62, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_63, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_64, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_65, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_31, c_0_28])).
cnf(c_0_66, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_xboole_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_67, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_68, plain, (k1_enumset1(X1,X1,X2)=k1_enumset1(X2,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_40]), c_0_40])).
cnf(c_0_69, plain, (r2_hidden(X1,X2)|r1_xboole_0(k1_tarski(X3),k1_tarski(X1))|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_57, c_0_32])).
cnf(c_0_70, negated_conjecture, (r2_hidden(esk8_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_71, plain, (X1=k1_tarski(X2)|~r2_hidden(X2,X1)|~r1_tarski(X1,k1_tarski(X2))), inference(spm,[status(thm)],[c_0_51, c_0_32])).
cnf(c_0_72, plain, (v1_xboole_0(X1)|r1_tarski(k1_tarski(esk2_1(X1)),X1)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_73, plain, (r2_hidden(X1,k1_tarski(esk2_1(k1_tarski(X1))))), inference(spm,[status(thm)],[c_0_61, c_0_36])).
cnf(c_0_74, plain, (r2_hidden(X1,X2)|X3!=k4_xboole_0(X4,k4_xboole_0(X4,X2))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_75, plain, (k4_xboole_0(X1,k1_tarski(X2))!=X1|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_76, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_77, plain, (r1_tarski(X1,X2)|~r1_xboole_0(X1,X3)|~r1_tarski(X1,k3_tarski(k1_enumset1(X2,X2,X3)))), inference(rw,[status(thm)],[c_0_66, c_0_55])).
cnf(c_0_78, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X2,X2,X1)))), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_79, negated_conjecture, (r2_hidden(X1,esk5_0)|r1_xboole_0(k1_tarski(esk8_0),k1_tarski(X1))), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_80, plain, (k1_tarski(esk2_1(k1_tarski(X1)))=k1_tarski(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73])]), c_0_48])).
cnf(c_0_81, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X3,k4_xboole_0(X3,X2)))), inference(er,[status(thm)],[c_0_74])).
cnf(c_0_82, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_75, c_0_76])).
fof(c_0_83, plain, ![X49]:(X49=k1_xboole_0|r2_hidden(esk4_1(X49),X49)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_84, plain, (r1_tarski(X1,X2)|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_85, negated_conjecture, (r2_hidden(esk2_1(k1_tarski(X1)),esk5_0)|r1_xboole_0(k1_tarski(esk8_0),k1_tarski(X1))), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_86, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X4!=k3_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_87, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_88, plain, (~r2_hidden(X1,k4_xboole_0(X2,X2))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_60]), c_0_82])).
cnf(c_0_89, plain, (X1=k1_xboole_0|r2_hidden(esk4_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_90, plain, (k4_xboole_0(X1,k1_tarski(X2))=X1|r2_hidden(esk2_1(k1_tarski(X2)),X1)), inference(spm,[status(thm)],[c_0_34, c_0_80])).
cnf(c_0_91, negated_conjecture, (r2_hidden(esk2_1(k1_tarski(esk8_0)),esk5_0)|r1_tarski(k1_tarski(esk8_0),X1)), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_92, plain, (r2_hidden(esk2_1(k1_tarski(X1)),X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_47, c_0_73])).
cnf(c_0_93, plain, (r2_hidden(X1,X4)|X4!=k4_xboole_0(X2,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_86, c_0_63])).
cnf(c_0_94, plain, (k4_xboole_0(k1_tarski(X1),X2)=k1_tarski(X1)|~r2_hidden(X1,k4_xboole_0(k1_tarski(X1),X2))), inference(spm,[status(thm)],[c_0_71, c_0_65])).
cnf(c_0_95, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_87])).
cnf(c_0_96, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_97, plain, (k4_xboole_0(k4_xboole_0(X1,X2),k1_tarski(X3))=k4_xboole_0(X1,X2)|~r2_hidden(esk2_1(k1_tarski(X3)),X2)), inference(spm,[status(thm)],[c_0_33, c_0_90])).
cnf(c_0_98, negated_conjecture, (r2_hidden(esk2_1(k1_tarski(esk8_0)),esk5_0)|r2_hidden(esk8_0,X1)), inference(spm,[status(thm)],[c_0_27, c_0_91])).
cnf(c_0_99, plain, (~r2_hidden(esk2_1(k1_tarski(X1)),X2)|~r2_hidden(X1,k4_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_33, c_0_92])).
cnf(c_0_100, plain, (r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_93])).
cnf(c_0_101, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_xboole_0|r2_hidden(esk4_1(k4_xboole_0(X1,k4_xboole_0(X1,X2))),X2)), inference(spm,[status(thm)],[c_0_81, c_0_89])).
cnf(c_0_102, plain, (k4_xboole_0(k1_tarski(X1),X2)=k1_tarski(X1)|r2_hidden(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_36])])).
cnf(c_0_103, plain, (k1_tarski(X1)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_96]), c_0_36])])).
cnf(c_0_104, plain, (r1_tarski(k4_xboole_0(k4_xboole_0(X1,X2),X3),X1)), inference(spm,[status(thm)],[c_0_31, c_0_65])).
cnf(c_0_105, plain, (r1_xboole_0(k4_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_42, c_0_28])).
cnf(c_0_106, plain, (~r2_hidden(esk2_1(k1_tarski(X1)),X2)|~r2_hidden(X3,k4_xboole_0(X4,X2))|~r2_hidden(X3,k1_tarski(X1))), inference(spm,[status(thm)],[c_0_33, c_0_97])).
cnf(c_0_107, negated_conjecture, (r2_hidden(esk2_1(k1_tarski(esk8_0)),esk5_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_98]), c_0_98])).
cnf(c_0_108, plain, (~r2_hidden(esk2_1(k1_tarski(X1)),k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_109, plain, (r2_hidden(X1,k4_xboole_0(k1_tarski(X1),X2))|r2_hidden(esk4_1(k1_tarski(X1)),X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_103])).
cnf(c_0_110, plain, (r1_tarski(k4_xboole_0(k4_xboole_0(k3_tarski(k1_enumset1(X1,X1,X2)),X3),X4),X1)|~r1_xboole_0(k4_xboole_0(k4_xboole_0(k3_tarski(k1_enumset1(X1,X1,X2)),X3),X4),X2)), inference(spm,[status(thm)],[c_0_77, c_0_104])).
cnf(c_0_111, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_105, c_0_60])).
cnf(c_0_112, negated_conjecture, (~r2_hidden(X1,k4_xboole_0(X2,esk5_0))|~r2_hidden(X1,k1_tarski(esk8_0))), inference(spm,[status(thm)],[c_0_106, c_0_107])).
cnf(c_0_113, plain, (r2_hidden(esk4_1(k1_tarski(X1)),X2)|~r2_hidden(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_80]), c_0_80]), c_0_36])])).
cnf(c_0_114, plain, (r1_tarski(k4_xboole_0(k4_xboole_0(k3_tarski(k1_enumset1(X1,X1,k1_xboole_0)),X2),X3),X1)), inference(spm,[status(thm)],[c_0_110, c_0_111])).
cnf(c_0_115, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_116, negated_conjecture, (~r2_hidden(esk4_1(k1_tarski(X1)),k1_tarski(esk8_0))|~r2_hidden(X1,k4_xboole_0(X2,esk5_0))), inference(spm,[status(thm)],[c_0_112, c_0_113])).
cnf(c_0_117, plain, (r1_tarski(k4_xboole_0(k3_tarski(k1_enumset1(X1,X1,k1_xboole_0)),X2),X1)), inference(spm,[status(thm)],[c_0_114, c_0_60])).
cnf(c_0_118, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_115])).
cnf(c_0_119, negated_conjecture, (~r2_hidden(esk2_1(k1_tarski(X1)),k4_xboole_0(X2,esk5_0))|~r2_hidden(esk4_1(k1_tarski(X1)),k1_tarski(esk8_0))), inference(spm,[status(thm)],[c_0_116, c_0_80])).
cnf(c_0_120, plain, (r2_hidden(esk2_1(k1_tarski(X1)),X2)|~r1_tarski(k1_tarski(X1),X2)), inference(spm,[status(thm)],[c_0_27, c_0_80])).
cnf(c_0_121, plain, (r1_tarski(k3_tarski(k1_enumset1(X1,X1,k1_xboole_0)),X1)), inference(spm,[status(thm)],[c_0_117, c_0_60])).
cnf(c_0_122, plain, (r2_hidden(esk3_3(X1,X2,X3),X1)|r2_hidden(esk3_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_123, plain, (r2_hidden(esk3_3(X1,X2,X3),X2)|r2_hidden(esk3_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_124, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|r2_hidden(esk4_1(k4_xboole_0(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_118, c_0_89])).
cnf(c_0_125, negated_conjecture, (~r2_hidden(esk4_1(k1_tarski(X1)),k1_tarski(esk8_0))|~r1_tarski(k1_tarski(X1),k4_xboole_0(X2,esk5_0))), inference(spm,[status(thm)],[c_0_119, c_0_120])).
cnf(c_0_126, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(esk2_1(k1_tarski(X1)),X2)), inference(spm,[status(thm)],[c_0_32, c_0_80])).
fof(c_0_127, plain, ![X59, X60, X61, X62, X65, X66, X67, X68, X69, X70, X72, X73]:((((r2_hidden(esk9_4(X59,X60,X61,X62),X59)|~r2_hidden(X62,X61)|X61!=k2_zfmisc_1(X59,X60))&((r2_hidden(esk10_4(X59,X60,X61,X62),X60)|~r2_hidden(X62,X61)|X61!=k2_zfmisc_1(X59,X60))&(X62=k4_tarski(esk9_4(X59,X60,X61,X62),esk10_4(X59,X60,X61,X62))|~r2_hidden(X62,X61)|X61!=k2_zfmisc_1(X59,X60))))&(~r2_hidden(X66,X59)|(~r2_hidden(X67,X60)|X65!=k4_tarski(X66,X67))|r2_hidden(X65,X61)|X61!=k2_zfmisc_1(X59,X60)))&((~r2_hidden(esk11_3(X68,X69,X70),X70)|(~r2_hidden(X72,X68)|(~r2_hidden(X73,X69)|esk11_3(X68,X69,X70)!=k4_tarski(X72,X73)))|X70=k2_zfmisc_1(X68,X69))&((r2_hidden(esk12_3(X68,X69,X70),X68)|r2_hidden(esk11_3(X68,X69,X70),X70)|X70=k2_zfmisc_1(X68,X69))&((r2_hidden(esk13_3(X68,X69,X70),X69)|r2_hidden(esk11_3(X68,X69,X70),X70)|X70=k2_zfmisc_1(X68,X69))&(esk11_3(X68,X69,X70)=k4_tarski(esk12_3(X68,X69,X70),esk13_3(X68,X69,X70))|r2_hidden(esk11_3(X68,X69,X70),X70)|X70=k2_zfmisc_1(X68,X69)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
cnf(c_0_128, plain, (k3_tarski(k1_enumset1(X1,X1,k1_xboole_0))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_121]), c_0_67])])).
fof(c_0_129, plain, ![X78, X79, X80]:(~r1_tarski(X78,X79)|~r1_xboole_0(X79,X80)|r1_xboole_0(X78,X80)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_xboole_1])])).
cnf(c_0_130, plain, (X3=k4_xboole_0(X1,k4_xboole_0(X1,X2))|r2_hidden(esk3_3(X1,X2,X3),X3)|r2_hidden(esk3_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_122, c_0_63])).
cnf(c_0_131, plain, (X3=k4_xboole_0(X1,k4_xboole_0(X1,X2))|r2_hidden(esk3_3(X1,X2,X3),X3)|r2_hidden(esk3_3(X1,X2,X3),X2)), inference(rw,[status(thm)],[c_0_123, c_0_63])).
cnf(c_0_132, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k1_xboole_0|~r2_hidden(esk4_1(k4_xboole_0(k4_xboole_0(X1,X2),X3)),X2)), inference(spm,[status(thm)],[c_0_33, c_0_124])).
cnf(c_0_133, negated_conjecture, (~r1_tarski(k1_tarski(esk8_0),k4_xboole_0(X1,esk5_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_89]), c_0_103])).
cnf(c_0_134, plain, (k4_xboole_0(X1,k1_tarski(X2))=X1|r1_tarski(k1_tarski(X2),X1)), inference(spm,[status(thm)],[c_0_126, c_0_90])).
cnf(c_0_135, plain, (X1=k4_tarski(esk9_4(X2,X3,X4,X1),esk10_4(X2,X3,X4,X1))|~r2_hidden(X1,X4)|X4!=k2_zfmisc_1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_136, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,k3_tarski(k1_enumset1(X2,X2,X3)))|~r1_xboole_0(k1_tarski(X1),X3)), inference(spm,[status(thm)],[c_0_77, c_0_32])).
cnf(c_0_137, plain, (k3_tarski(k1_enumset1(k1_xboole_0,k1_xboole_0,X1))=X1), inference(spm,[status(thm)],[c_0_128, c_0_68])).
cnf(c_0_138, plain, (~r1_tarski(k1_tarski(X1),k1_xboole_0)), inference(spm,[status(thm)],[c_0_82, c_0_120])).
cnf(c_0_139, plain, (r1_xboole_0(X1,X3)|~r1_tarski(X1,X2)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_140, plain, (X1=k4_xboole_0(k4_xboole_0(X2,X3),k4_xboole_0(k4_xboole_0(X2,X3),X4))|r2_hidden(esk3_3(k4_xboole_0(X2,X3),X4,X1),X1)|~r2_hidden(esk3_3(k4_xboole_0(X2,X3),X4,X1),X3)), inference(spm,[status(thm)],[c_0_33, c_0_130])).
cnf(c_0_141, plain, (k1_tarski(esk2_1(k4_xboole_0(X1,k4_xboole_0(X1,X2))))=k4_xboole_0(X1,k4_xboole_0(X1,X2))|r2_hidden(esk3_3(X1,X2,k1_tarski(X3)),k1_tarski(X3))|r2_hidden(esk3_3(X1,X2,k1_tarski(X3)),X2)), inference(spm,[status(thm)],[c_0_80, c_0_131])).
cnf(c_0_142, plain, (k4_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(k4_xboole_0(X1,X2),X2))=k1_xboole_0), inference(spm,[status(thm)],[c_0_132, c_0_101])).
cnf(c_0_143, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=X2|r2_hidden(esk3_3(X1,X2,X2),X2)), inference(ef,[status(thm)],[c_0_131])).
cnf(c_0_144, negated_conjecture, (k4_xboole_0(k4_xboole_0(X1,esk5_0),k1_tarski(esk8_0))=k4_xboole_0(X1,esk5_0)), inference(spm,[status(thm)],[c_0_133, c_0_134])).
cnf(c_0_145, negated_conjecture, (~r2_hidden(X1,esk6_0)|~r2_hidden(X2,esk7_0)|esk8_0!=k4_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_146, plain, (k4_tarski(esk9_4(X1,X2,k2_zfmisc_1(X1,X2),X3),esk10_4(X1,X2,k2_zfmisc_1(X1,X2),X3))=X3|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_135])).
cnf(c_0_147, plain, (r2_hidden(esk10_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_148, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(k1_tarski(X1),X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_137]), c_0_138])).
cnf(c_0_149, plain, (r1_xboole_0(k1_tarski(X1),X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_139, c_0_32])).
cnf(c_0_150, plain, (r2_hidden(esk3_3(k4_xboole_0(X1,X2),X2,k1_tarski(X3)),k1_tarski(X3))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_141]), c_0_142]), c_0_142]), c_0_142]), c_0_103]), c_0_103])).
cnf(c_0_151, plain, (r2_hidden(X1,k1_tarski(X2))|~r2_hidden(X2,k1_tarski(X1))), inference(spm,[status(thm)],[c_0_75, c_0_102])).
cnf(c_0_152, negated_conjecture, (r2_hidden(esk3_3(k4_xboole_0(X1,esk5_0),k1_tarski(esk8_0),k1_tarski(esk8_0)),k1_tarski(esk8_0))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144]), c_0_96]), c_0_103])).
cnf(c_0_153, negated_conjecture, (~r2_hidden(esk10_4(X1,X2,k2_zfmisc_1(X1,X2),esk8_0),esk7_0)|~r2_hidden(esk9_4(X1,X2,k2_zfmisc_1(X1,X2),esk8_0),esk6_0)|~r2_hidden(esk8_0,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146])])).
cnf(c_0_154, plain, (r2_hidden(esk10_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X2)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_147])).
cnf(c_0_155, plain, (r2_hidden(esk9_4(X1,X2,X3,X4),X1)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_156, plain, (~r2_hidden(esk2_1(k1_tarski(X1)),X2)|~r1_xboole_0(k1_tarski(X1),X2)), inference(spm,[status(thm)],[c_0_148, c_0_80])).
cnf(c_0_157, negated_conjecture, (r1_xboole_0(k1_tarski(esk8_0),X1)|~r1_xboole_0(esk5_0,X1)), inference(spm,[status(thm)],[c_0_149, c_0_70])).
cnf(c_0_158, negated_conjecture, (r2_hidden(esk3_3(k4_xboole_0(X1,esk5_0),k1_tarski(esk8_0),k1_tarski(X2)),k1_tarski(X2))), inference(spm,[status(thm)],[c_0_150, c_0_144])).
cnf(c_0_159, plain, (k1_tarski(X1)=k1_tarski(X2)|~r2_hidden(X2,k1_tarski(X1))|~r2_hidden(X1,k1_tarski(X2))), inference(spm,[status(thm)],[c_0_71, c_0_32])).
cnf(c_0_160, negated_conjecture, (r2_hidden(esk8_0,k1_tarski(esk3_3(k4_xboole_0(X1,esk5_0),k1_tarski(esk8_0),k1_tarski(esk8_0))))), inference(spm,[status(thm)],[c_0_151, c_0_152])).
cnf(c_0_161, negated_conjecture, (~r2_hidden(esk9_4(X1,esk7_0,k2_zfmisc_1(X1,esk7_0),esk8_0),esk6_0)|~r2_hidden(esk8_0,k2_zfmisc_1(X1,esk7_0))), inference(spm,[status(thm)],[c_0_153, c_0_154])).
cnf(c_0_162, plain, (r2_hidden(esk9_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X1)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_155])).
cnf(c_0_163, negated_conjecture, (r1_tarski(esk5_0,k2_zfmisc_1(esk6_0,esk7_0))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_164, negated_conjecture, (~r2_hidden(esk2_1(k1_tarski(esk8_0)),X1)|~r1_xboole_0(esk5_0,X1)), inference(spm,[status(thm)],[c_0_156, c_0_157])).
cnf(c_0_165, negated_conjecture, (r2_hidden(X1,k1_tarski(esk3_3(k4_xboole_0(X2,esk5_0),k1_tarski(esk8_0),k1_tarski(X1))))), inference(spm,[status(thm)],[c_0_151, c_0_158])).
cnf(c_0_166, negated_conjecture, (k1_tarski(esk3_3(k4_xboole_0(X1,esk5_0),k1_tarski(esk8_0),k1_tarski(esk8_0)))=k1_tarski(esk8_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_160]), c_0_152])])).
cnf(c_0_167, negated_conjecture, (~r2_hidden(esk8_0,k2_zfmisc_1(esk6_0,esk7_0))), inference(spm,[status(thm)],[c_0_161, c_0_162])).
cnf(c_0_168, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(esk6_0,esk7_0))|r1_xboole_0(esk5_0,k1_tarski(X1))), inference(spm,[status(thm)],[c_0_57, c_0_163])).
cnf(c_0_169, negated_conjecture, (~r1_xboole_0(esk5_0,k1_tarski(esk8_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_165]), c_0_80]), c_0_166])).
cnf(c_0_170, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_168]), c_0_169]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 171
# Proof object clause steps            : 132
# Proof object formula steps           : 39
# Proof object conjectures             : 30
# Proof object clause conjectures      : 27
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 33
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 82
# Proof object simplifying inferences  : 50
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 44
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 41
# Processed clauses                    : 10830
# ...of these trivial                  : 127
# ...subsumed                          : 9225
# ...remaining for further processing  : 1478
# Other redundant clauses eliminated   : 291
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 49
# Backward-rewritten                   : 59
# Generated clauses                    : 245151
# ...of the previous two non-trivial   : 223960
# Contextual simplify-reflections      : 20
# Paramodulations                      : 244511
# Factorizations                       : 350
# Equation resolutions                 : 291
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1358
#    Positive orientable unit clauses  : 159
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 37
#    Non-unit-clauses                  : 1161
# Current number of unprocessed clauses: 212691
# ...number of literals in the above   : 960656
# Current number of archived formulas  : 0
# Current number of archived clauses   : 111
# Clause-clause subsumption calls (NU) : 275111
# Rec. Clause-clause subsumption calls : 147252
# Non-unit clause-clause subsumptions  : 6324
# Unit Clause-clause subsumption calls : 8266
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 364
# BW rewrite match successes           : 51
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4777344

# -------------------------------------------------
# User time                : 3.637 s
# System time              : 0.093 s
# Total time               : 3.730 s
# Maximum resident set size: 164252 pages
