# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', t2_xboole_1)).
fof(d13_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2, X3]:(X3=k10_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,k1_relat_1(X1))&r2_hidden(k1_funct_1(X1,X4),X2))))), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', d13_funct_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', d3_tarski)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', l32_xboole_1)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', redefinition_k6_subset_1)).
fof(t3_xboole_1, axiom, ![X1]:(r1_tarski(X1,k1_xboole_0)=>X1=k1_xboole_0), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', t3_xboole_1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', t12_setfam_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', t70_enumset1)).
fof(t138_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>k10_relat_1(X3,k6_subset_1(X1,X2))=k6_subset_1(k10_relat_1(X3,X1),k10_relat_1(X3,X2))), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', t138_funct_1)).
fof(t169_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k10_relat_1(X1,k2_relat_1(X1))=k1_relat_1(X1)), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', t169_relat_1)).
fof(t71_relat_1, axiom, ![X1]:(k1_relat_1(k6_relat_1(X1))=X1&k2_relat_1(k6_relat_1(X1))=X1), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', t71_relat_1)).
fof(fc3_funct_1, axiom, ![X1]:(v1_relat_1(k6_relat_1(X1))&v1_funct_1(k6_relat_1(X1))), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', fc3_funct_1)).
fof(t167_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k10_relat_1(X2,X1),k1_relat_1(X2))), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', t167_relat_1)).
fof(t175_funct_2, conjecture, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2, X3]:(r1_tarski(k10_relat_1(X1,X3),X2)=>k10_relat_1(X1,X3)=k10_relat_1(k7_relat_1(X1,X2),X3))), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', t175_funct_2)).
fof(t139_funct_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>k10_relat_1(k7_relat_1(X3,X1),X2)=k3_xboole_0(X1,k10_relat_1(X3,X2))), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', t139_funct_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', t48_xboole_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', commutativity_k2_tarski)).
fof(t147_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r1_tarski(X1,k2_relat_1(X2))=>k9_relat_1(X2,k10_relat_1(X2,X1))=X1)), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', t147_funct_1)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', t1_xboole_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', t36_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/54c4b53c6b594759aefd65b90abc7fc5.E_input', t3_boole)).
fof(c_0_22, plain, ![X5, X6]:(~r2_hidden(X5,X6)|~r1_tarski(X6,X5)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_23, plain, ![X7]:r1_tarski(k1_xboole_0,X7), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_24, plain, ![X8, X9, X10, X11, X12, X13, X14]:((((r2_hidden(X11,k1_relat_1(X8))|~r2_hidden(X11,X10)|X10!=k10_relat_1(X8,X9)|(~v1_relat_1(X8)|~v1_funct_1(X8)))&(r2_hidden(k1_funct_1(X8,X11),X9)|~r2_hidden(X11,X10)|X10!=k10_relat_1(X8,X9)|(~v1_relat_1(X8)|~v1_funct_1(X8))))&(~r2_hidden(X12,k1_relat_1(X8))|~r2_hidden(k1_funct_1(X8,X12),X9)|r2_hidden(X12,X10)|X10!=k10_relat_1(X8,X9)|(~v1_relat_1(X8)|~v1_funct_1(X8))))&((~r2_hidden(esk1_3(X8,X13,X14),X14)|(~r2_hidden(esk1_3(X8,X13,X14),k1_relat_1(X8))|~r2_hidden(k1_funct_1(X8,esk1_3(X8,X13,X14)),X13))|X14=k10_relat_1(X8,X13)|(~v1_relat_1(X8)|~v1_funct_1(X8)))&((r2_hidden(esk1_3(X8,X13,X14),k1_relat_1(X8))|r2_hidden(esk1_3(X8,X13,X14),X14)|X14=k10_relat_1(X8,X13)|(~v1_relat_1(X8)|~v1_funct_1(X8)))&(r2_hidden(k1_funct_1(X8,esk1_3(X8,X13,X14)),X13)|r2_hidden(esk1_3(X8,X13,X14),X14)|X14=k10_relat_1(X8,X13)|(~v1_relat_1(X8)|~v1_funct_1(X8)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d13_funct_1])])])])])])).
cnf(c_0_25, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27, plain, (r2_hidden(k1_funct_1(X1,X2),X3)|~r2_hidden(X2,X4)|X4!=k10_relat_1(X1,X3)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_28, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_29, plain, (r2_hidden(k1_funct_1(X1,X2),X3)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X2,k10_relat_1(X1,X3))), inference(er,[status(thm)],[c_0_27])).
fof(c_0_30, plain, ![X16, X17, X18, X19, X20]:((~r1_tarski(X16,X17)|(~r2_hidden(X18,X16)|r2_hidden(X18,X17)))&((r2_hidden(esk2_2(X19,X20),X19)|r1_tarski(X19,X20))&(~r2_hidden(esk2_2(X19,X20),X20)|r1_tarski(X19,X20)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_31, plain, ![X32, X33]:((k4_xboole_0(X32,X33)!=k1_xboole_0|r1_tarski(X32,X33))&(~r1_tarski(X32,X33)|k4_xboole_0(X32,X33)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
fof(c_0_32, plain, ![X25, X26]:k6_subset_1(X25,X26)=k4_xboole_0(X25,X26), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
fof(c_0_33, plain, ![X27]:(~r1_tarski(X27,k1_xboole_0)|X27=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_xboole_1])])).
cnf(c_0_34, plain, (~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X2,k10_relat_1(X1,k1_xboole_0))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_36, plain, ![X28, X29]:k1_setfam_1(k2_tarski(X28,X29))=k3_xboole_0(X28,X29), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_37, plain, ![X30, X31]:k1_enumset1(X30,X30,X31)=k2_tarski(X30,X31), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_38, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_39, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_40, plain, ![X22, X23, X24]:(~v1_relat_1(X24)|~v1_funct_1(X24)|k10_relat_1(X24,k6_subset_1(X22,X23))=k6_subset_1(k10_relat_1(X24,X22),k10_relat_1(X24,X23))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t138_funct_1])])).
fof(c_0_41, plain, ![X34]:(~v1_relat_1(X34)|k10_relat_1(X34,k2_relat_1(X34))=k1_relat_1(X34)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t169_relat_1])])).
fof(c_0_42, plain, ![X35]:(k1_relat_1(k6_relat_1(X35))=X35&k2_relat_1(k6_relat_1(X35))=X35), inference(variable_rename,[status(thm)],[t71_relat_1])).
fof(c_0_43, plain, ![X36]:(v1_relat_1(k6_relat_1(X36))&v1_funct_1(k6_relat_1(X36))), inference(variable_rename,[status(thm)],[fc3_funct_1])).
fof(c_0_44, plain, ![X37, X38]:(~v1_relat_1(X38)|r1_tarski(k10_relat_1(X38,X37),k1_relat_1(X38))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t167_relat_1])])).
cnf(c_0_45, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_46, plain, (r1_tarski(k10_relat_1(X1,k1_xboole_0),X2)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
fof(c_0_47, negated_conjecture, ~(![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2, X3]:(r1_tarski(k10_relat_1(X1,X3),X2)=>k10_relat_1(X1,X3)=k10_relat_1(k7_relat_1(X1,X2),X3)))), inference(assume_negation,[status(cth)],[t175_funct_2])).
fof(c_0_48, plain, ![X42, X43, X44]:(~v1_relat_1(X44)|k10_relat_1(k7_relat_1(X44,X42),X43)=k3_xboole_0(X42,k10_relat_1(X44,X43))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t139_funct_1])])).
cnf(c_0_49, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_50, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_51, plain, ![X45, X46]:k4_xboole_0(X45,k4_xboole_0(X45,X46))=k3_xboole_0(X45,X46), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_52, plain, ![X47, X48]:k2_tarski(X47,X48)=k2_tarski(X48,X47), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_53, plain, (k6_subset_1(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_54, plain, (k10_relat_1(X1,k6_subset_1(X2,X3))=k6_subset_1(k10_relat_1(X1,X2),k10_relat_1(X1,X3))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_55, plain, (k10_relat_1(X1,k2_relat_1(X1))=k1_relat_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_56, plain, (k2_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_57, plain, (k1_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_58, plain, (v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_59, plain, (r1_tarski(k10_relat_1(X1,X2),k1_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_60, plain, ![X49, X50]:(~v1_relat_1(X50)|~v1_funct_1(X50)|(~r1_tarski(X49,k2_relat_1(X50))|k9_relat_1(X50,k10_relat_1(X50,X49))=X49)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t147_funct_1])])).
cnf(c_0_61, plain, (k10_relat_1(X1,k1_xboole_0)=k1_xboole_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_62, plain, (v1_funct_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_63, plain, ![X51, X52, X53]:(~r1_tarski(X51,X52)|~r1_tarski(X52,X53)|r1_tarski(X51,X53)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_64, negated_conjecture, ((v1_relat_1(esk3_0)&v1_funct_1(esk3_0))&(r1_tarski(k10_relat_1(esk3_0,esk5_0),esk4_0)&k10_relat_1(esk3_0,esk5_0)!=k10_relat_1(k7_relat_1(esk3_0,esk4_0),esk5_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_47])])])).
fof(c_0_65, plain, ![X54, X55]:r1_tarski(k4_xboole_0(X54,X55),X54), inference(variable_rename,[status(thm)],[t36_xboole_1])).
cnf(c_0_66, plain, (k10_relat_1(k7_relat_1(X1,X2),X3)=k3_xboole_0(X2,k10_relat_1(X1,X3))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_67, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_68, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_69, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_70, plain, (k10_relat_1(X1,k6_subset_1(X2,X3))=k1_xboole_0|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_tarski(k10_relat_1(X1,X2),k10_relat_1(X1,X3))), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_71, plain, (k10_relat_1(k6_relat_1(X1),X1)=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57]), c_0_58])])).
cnf(c_0_72, plain, (r1_tarski(k10_relat_1(k6_relat_1(X1),X2),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_57]), c_0_58])])).
cnf(c_0_73, plain, (k9_relat_1(X1,k10_relat_1(X1,X2))=X2|~v1_relat_1(X1)|~v1_funct_1(X1)|~r1_tarski(X2,k2_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_74, plain, (k10_relat_1(k6_relat_1(X1),k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_58])])).
cnf(c_0_75, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_76, negated_conjecture, (r1_tarski(k10_relat_1(esk3_0,esk5_0),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_77, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_78, plain, (k10_relat_1(k7_relat_1(X1,X2),X3)=k1_setfam_1(k1_enumset1(X2,X2,k10_relat_1(X1,X3)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_79, plain, (k6_subset_1(X1,k6_subset_1(X1,X2))=k1_setfam_1(k1_enumset1(X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_39]), c_0_39]), c_0_67])).
cnf(c_0_80, plain, (k1_enumset1(X1,X1,X2)=k1_enumset1(X2,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_50]), c_0_50])).
cnf(c_0_81, plain, (k10_relat_1(k6_relat_1(X1),k6_subset_1(X2,X1))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_62]), c_0_58]), c_0_72])])).
cnf(c_0_82, plain, (k9_relat_1(k6_relat_1(X1),k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_62]), c_0_58]), c_0_56]), c_0_26])])).
cnf(c_0_83, negated_conjecture, (r1_tarski(X1,esk4_0)|~r1_tarski(X1,k10_relat_1(esk3_0,esk5_0))), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_84, plain, (r1_tarski(k6_subset_1(X1,X2),X1)), inference(rw,[status(thm)],[c_0_77, c_0_39])).
fof(c_0_85, plain, ![X56]:k4_xboole_0(X56,k1_xboole_0)=X56, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_86, plain, (k6_subset_1(X1,k6_subset_1(X1,k10_relat_1(X2,X3)))=k10_relat_1(k7_relat_1(X2,X1),X3)|~v1_relat_1(X2)), inference(rw,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_87, plain, (k6_subset_1(X1,k6_subset_1(X1,X2))=k6_subset_1(X2,k6_subset_1(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_79])).
cnf(c_0_88, plain, (k6_subset_1(X1,X2)=k1_xboole_0|~r1_tarski(k6_subset_1(X1,X2),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_81]), c_0_82]), c_0_62]), c_0_58]), c_0_56])])).
cnf(c_0_89, negated_conjecture, (r1_tarski(k6_subset_1(k10_relat_1(esk3_0,esk5_0),X1),esk4_0)), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_90, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_91, plain, (k6_subset_1(k10_relat_1(X1,X2),k6_subset_1(k10_relat_1(X1,X2),X3))=k10_relat_1(k7_relat_1(X1,X3),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_92, negated_conjecture, (k6_subset_1(k10_relat_1(esk3_0,esk5_0),esk4_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_93, plain, (k6_subset_1(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_90, c_0_39])).
cnf(c_0_94, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_95, negated_conjecture, (k10_relat_1(esk3_0,esk5_0)!=k10_relat_1(k7_relat_1(esk3_0,esk4_0),esk5_0)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_96, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_93]), c_0_94])]), c_0_95]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 97
# Proof object clause steps            : 52
# Proof object formula steps           : 45
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 22
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 38
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 32
# Processed clauses                    : 11095
# ...of these trivial                  : 60
# ...subsumed                          : 9865
# ...remaining for further processing  : 1170
# Other redundant clauses eliminated   : 53
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 63
# Backward-rewritten                   : 56
# Generated clauses                    : 74858
# ...of the previous two non-trivial   : 66352
# Contextual simplify-reflections      : 8
# Paramodulations                      : 74801
# Factorizations                       : 4
# Equation resolutions                 : 53
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1016
#    Positive orientable unit clauses  : 103
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 79
#    Non-unit-clauses                  : 832
# Current number of unprocessed clauses: 54884
# ...number of literals in the above   : 219998
# Current number of archived formulas  : 0
# Current number of archived clauses   : 154
# Clause-clause subsumption calls (NU) : 179079
# Rec. Clause-clause subsumption calls : 56409
# Non-unit clause-clause subsumptions  : 4486
# Unit Clause-clause subsumption calls : 9979
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 559
# BW rewrite match successes           : 40
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1388635

# -------------------------------------------------
# User time                : 0.804 s
# System time              : 0.020 s
# Total time               : 0.824 s
# Maximum resident set size: 53068 pages
