# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t1_yellow_1, axiom, ![X1]:(u1_struct_0(k2_yellow_1(X1))=X1&u1_orders_2(k2_yellow_1(X1))=k1_yellow_1(X1)), file('data/MPTP/proofs/5522b7f34ce64744b5182eff0c33e128.E_input', t1_yellow_1)).
fof(redefinition_k1_yellow_1, axiom, ![X1]:k1_yellow_1(X1)=k1_wellord2(X1), file('data/MPTP/proofs/5522b7f34ce64744b5182eff0c33e128.E_input', redefinition_k1_yellow_1)).
fof(d1_wellord2, axiom, ![X1, X2]:(v1_relat_1(X2)=>(X2=k1_wellord2(X1)<=>(k3_relat_1(X2)=X1&![X3, X4]:((r2_hidden(X3,X1)&r2_hidden(X4,X1))=>(r2_hidden(k4_tarski(X3,X4),X2)<=>r1_tarski(X3,X4)))))), file('data/MPTP/proofs/5522b7f34ce64744b5182eff0c33e128.E_input', d1_wellord2)).
fof(dt_k1_wellord2, axiom, ![X1]:v1_relat_1(k1_wellord2(X1)), file('data/MPTP/proofs/5522b7f34ce64744b5182eff0c33e128.E_input', dt_k1_wellord2)).
fof(t3_yellow_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r3_orders_2(k2_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3))))), file('data/MPTP/proofs/5522b7f34ce64744b5182eff0c33e128.E_input', t3_yellow_1)).
fof(d9_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)<=>r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1)))))), file('data/MPTP/proofs/5522b7f34ce64744b5182eff0c33e128.E_input', d9_orders_2)).
fof(dt_k2_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&l1_orders_2(k2_yellow_1(X1))), file('data/MPTP/proofs/5522b7f34ce64744b5182eff0c33e128.E_input', dt_k2_yellow_1)).
fof(fc1_struct_0, axiom, ![X1]:((v2_struct_0(X1)&l1_struct_0(X1))=>v1_xboole_0(u1_struct_0(X1))), file('data/MPTP/proofs/5522b7f34ce64744b5182eff0c33e128.E_input', fc1_struct_0)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/5522b7f34ce64744b5182eff0c33e128.E_input', redefinition_r3_orders_2)).
fof(fc5_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&(v3_orders_2(k2_yellow_1(X1))&(v4_orders_2(k2_yellow_1(X1))&v5_orders_2(k2_yellow_1(X1))))), file('data/MPTP/proofs/5522b7f34ce64744b5182eff0c33e128.E_input', fc5_yellow_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/5522b7f34ce64744b5182eff0c33e128.E_input', t2_subset)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/5522b7f34ce64744b5182eff0c33e128.E_input', dt_l1_orders_2)).
fof(c_0_12, plain, ![X12]:(u1_struct_0(k2_yellow_1(X12))=X12&u1_orders_2(k2_yellow_1(X12))=k1_yellow_1(X12)), inference(variable_rename,[status(thm)],[t1_yellow_1])).
fof(c_0_13, plain, ![X11]:k1_yellow_1(X11)=k1_wellord2(X11), inference(variable_rename,[status(thm)],[redefinition_k1_yellow_1])).
fof(c_0_14, plain, ![X5, X6, X7, X8]:(((k3_relat_1(X6)=X5|X6!=k1_wellord2(X5)|~v1_relat_1(X6))&((~r2_hidden(k4_tarski(X7,X8),X6)|r1_tarski(X7,X8)|(~r2_hidden(X7,X5)|~r2_hidden(X8,X5))|X6!=k1_wellord2(X5)|~v1_relat_1(X6))&(~r1_tarski(X7,X8)|r2_hidden(k4_tarski(X7,X8),X6)|(~r2_hidden(X7,X5)|~r2_hidden(X8,X5))|X6!=k1_wellord2(X5)|~v1_relat_1(X6))))&(((r2_hidden(esk1_2(X5,X6),X5)|k3_relat_1(X6)!=X5|X6=k1_wellord2(X5)|~v1_relat_1(X6))&(r2_hidden(esk2_2(X5,X6),X5)|k3_relat_1(X6)!=X5|X6=k1_wellord2(X5)|~v1_relat_1(X6)))&((~r2_hidden(k4_tarski(esk1_2(X5,X6),esk2_2(X5,X6)),X6)|~r1_tarski(esk1_2(X5,X6),esk2_2(X5,X6))|k3_relat_1(X6)!=X5|X6=k1_wellord2(X5)|~v1_relat_1(X6))&(r2_hidden(k4_tarski(esk1_2(X5,X6),esk2_2(X5,X6)),X6)|r1_tarski(esk1_2(X5,X6),esk2_2(X5,X6))|k3_relat_1(X6)!=X5|X6=k1_wellord2(X5)|~v1_relat_1(X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_wellord2])])])])])).
cnf(c_0_15, plain, (u1_orders_2(k2_yellow_1(X1))=k1_yellow_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, plain, (k1_yellow_1(X1)=k1_wellord2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_17, plain, ![X13]:v1_relat_1(k1_wellord2(X13)), inference(variable_rename,[status(thm)],[dt_k1_wellord2])).
cnf(c_0_18, plain, (r1_tarski(X1,X2)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,X4)|~r2_hidden(X2,X4)|X3!=k1_wellord2(X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (u1_orders_2(k2_yellow_1(X1))=k1_wellord2(X1)), inference(rw,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_20, plain, (v1_relat_1(k1_wellord2(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_21, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r3_orders_2(k2_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3)))))), inference(assume_negation,[status(cth)],[t3_yellow_1])).
cnf(c_0_22, plain, (r1_tarski(X1,X2)|X3!=u1_orders_2(k2_yellow_1(X4))|~v1_relat_1(X3)|~r2_hidden(X2,X4)|~r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)), inference(rw,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_23, plain, (v1_relat_1(u1_orders_2(k2_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_20, c_0_19])).
fof(c_0_24, plain, ![X17, X18, X19]:((~r1_orders_2(X17,X18,X19)|r2_hidden(k4_tarski(X18,X19),u1_orders_2(X17))|~m1_subset_1(X19,u1_struct_0(X17))|~m1_subset_1(X18,u1_struct_0(X17))|~l1_orders_2(X17))&(~r2_hidden(k4_tarski(X18,X19),u1_orders_2(X17))|r1_orders_2(X17,X18,X19)|~m1_subset_1(X19,u1_struct_0(X17))|~m1_subset_1(X18,u1_struct_0(X17))|~l1_orders_2(X17))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_orders_2])])])])).
fof(c_0_25, plain, ![X20]:(v1_orders_2(k2_yellow_1(X20))&l1_orders_2(k2_yellow_1(X20))), inference(variable_rename,[status(thm)],[dt_k2_yellow_1])).
fof(c_0_26, negated_conjecture, (~v1_xboole_0(esk3_0)&(m1_subset_1(esk4_0,u1_struct_0(k2_yellow_1(esk3_0)))&(m1_subset_1(esk5_0,u1_struct_0(k2_yellow_1(esk3_0)))&((~r3_orders_2(k2_yellow_1(esk3_0),esk4_0,esk5_0)|~r1_tarski(esk4_0,esk5_0))&(r3_orders_2(k2_yellow_1(esk3_0),esk4_0,esk5_0)|r1_tarski(esk4_0,esk5_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_21])])])])).
fof(c_0_27, plain, ![X21]:(~v2_struct_0(X21)|~l1_struct_0(X21)|v1_xboole_0(u1_struct_0(X21))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_struct_0])])).
cnf(c_0_28, plain, (r1_tarski(X1,X2)|~r2_hidden(k4_tarski(X1,X2),u1_orders_2(k2_yellow_1(X3)))|~r2_hidden(X2,X3)|~r2_hidden(X1,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_22]), c_0_23])])).
cnf(c_0_29, plain, (r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, plain, (l1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, plain, (u1_struct_0(k2_yellow_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_32, plain, ![X22, X23, X24]:((~r3_orders_2(X22,X23,X24)|r1_orders_2(X22,X23,X24)|(v2_struct_0(X22)|(~v3_orders_2(X22)|~l1_orders_2(X22))|(~m1_subset_1(X23,u1_struct_0(X22))|~m1_subset_1(X24,u1_struct_0(X22)))))&(~r1_orders_2(X22,X23,X24)|r3_orders_2(X22,X23,X24)|(v2_struct_0(X22)|(~v3_orders_2(X22)|~l1_orders_2(X22))|(~m1_subset_1(X23,u1_struct_0(X22))|~m1_subset_1(X24,u1_struct_0(X22)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
fof(c_0_33, plain, ![X25]:(v1_orders_2(k2_yellow_1(X25))&(v3_orders_2(k2_yellow_1(X25))&(v4_orders_2(k2_yellow_1(X25))&v5_orders_2(k2_yellow_1(X25))))), inference(variable_rename,[status(thm)],[fc5_yellow_1])).
fof(c_0_34, plain, ![X26, X27]:(~m1_subset_1(X26,X27)|(v1_xboole_0(X27)|r2_hidden(X26,X27))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_35, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(k2_yellow_1(esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(k2_yellow_1(esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_37, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_38, plain, ![X28]:(~l1_orders_2(X28)|l1_struct_0(X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_39, plain, (r1_tarski(X1,X2)|~r1_orders_2(k2_yellow_1(X3),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3)|~r2_hidden(X2,X3)|~r2_hidden(X1,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30]), c_0_31]), c_0_31])])).
cnf(c_0_40, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_41, plain, (v3_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43, negated_conjecture, (m1_subset_1(esk5_0,esk3_0)), inference(rw,[status(thm)],[c_0_35, c_0_31])).
cnf(c_0_44, negated_conjecture, (~v1_xboole_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_45, negated_conjecture, (m1_subset_1(esk4_0,esk3_0)), inference(rw,[status(thm)],[c_0_36, c_0_31])).
cnf(c_0_46, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r1_tarski(X1,X2)|~r2_hidden(X1,X4)|~r2_hidden(X2,X4)|X3!=k1_wellord2(X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_47, plain, (v1_xboole_0(X1)|~l1_struct_0(k2_yellow_1(X1))|~v2_struct_0(k2_yellow_1(X1))), inference(spm,[status(thm)],[c_0_37, c_0_31])).
cnf(c_0_48, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_49, plain, (v2_struct_0(k2_yellow_1(X1))|r1_tarski(X2,X3)|~r3_orders_2(k2_yellow_1(X1),X2,X3)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~r2_hidden(X3,X1)|~r2_hidden(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41]), c_0_30]), c_0_31]), c_0_31])])).
cnf(c_0_50, negated_conjecture, (r3_orders_2(k2_yellow_1(esk3_0),esk4_0,esk5_0)|r1_tarski(esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_51, negated_conjecture, (r2_hidden(esk5_0,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk4_0,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_45]), c_0_44])).
cnf(c_0_53, plain, (r2_hidden(k4_tarski(X1,X2),X3)|X3!=u1_orders_2(k2_yellow_1(X4))|~v1_relat_1(X3)|~r2_hidden(X2,X4)|~r2_hidden(X1,X4)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_46, c_0_19])).
cnf(c_0_54, plain, (v1_xboole_0(X1)|~v2_struct_0(k2_yellow_1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_30])])).
cnf(c_0_55, negated_conjecture, (v2_struct_0(k2_yellow_1(esk3_0))|r1_tarski(esk4_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_43]), c_0_45]), c_0_51]), c_0_52])])).
cnf(c_0_56, plain, (r1_orders_2(X3,X1,X2)|~r2_hidden(k4_tarski(X1,X2),u1_orders_2(X3))|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_57, plain, (r2_hidden(k4_tarski(X1,X2),u1_orders_2(k2_yellow_1(X3)))|~r1_tarski(X1,X2)|~r2_hidden(X2,X3)|~r2_hidden(X1,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_53]), c_0_23])])).
cnf(c_0_58, negated_conjecture, (~r3_orders_2(k2_yellow_1(esk3_0),esk4_0,esk5_0)|~r1_tarski(esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_59, negated_conjecture, (r1_tarski(esk4_0,esk5_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_44])).
cnf(c_0_60, plain, (r3_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_61, plain, (r1_orders_2(k2_yellow_1(X1),X2,X3)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~r1_tarski(X2,X3)|~r2_hidden(X3,X1)|~r2_hidden(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_30]), c_0_31]), c_0_31])])).
cnf(c_0_62, negated_conjecture, (~r3_orders_2(k2_yellow_1(esk3_0),esk4_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_59])])).
cnf(c_0_63, plain, (v2_struct_0(k2_yellow_1(X1))|r3_orders_2(k2_yellow_1(X1),X2,X3)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~r1_tarski(X2,X3)|~r2_hidden(X3,X1)|~r2_hidden(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_41]), c_0_30]), c_0_31]), c_0_31])])).
cnf(c_0_64, negated_conjecture, (v2_struct_0(k2_yellow_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_43]), c_0_45]), c_0_59]), c_0_51]), c_0_52])])).
cnf(c_0_65, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_64]), c_0_44]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 66
# Proof object clause steps            : 41
# Proof object formula steps           : 25
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 49
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 27
# Processed clauses                    : 74
# ...of these trivial                  : 1
# ...subsumed                          : 0
# ...remaining for further processing  : 73
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 5
# Generated clauses                    : 30
# ...of the previous two non-trivial   : 23
# Contextual simplify-reflections      : 0
# Paramodulations                      : 23
# Factorizations                       : 0
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 34
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 18
# Current number of unprocessed clauses: 2
# ...number of literals in the above   : 14
# Current number of archived formulas  : 0
# Current number of archived clauses   : 34
# Clause-clause subsumption calls (NU) : 238
# Rec. Clause-clause subsumption calls : 72
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3009

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 29024 pages
