# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/5527c7ecc9b943a8838ba6cba9a3b4d3.E_input', symmetry_r1_xboole_0)).
fof(t79_xboole_1, axiom, ![X1, X2]:r1_xboole_0(k4_xboole_0(X1,X2),X2), file('data/MPTP/proofs/5527c7ecc9b943a8838ba6cba9a3b4d3.E_input', t79_xboole_1)).
fof(d14_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(X3=k10_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(k4_tarski(X4,X5),X1)&r2_hidden(X5,X2))))), file('data/MPTP/proofs/5527c7ecc9b943a8838ba6cba9a3b4d3.E_input', d14_relat_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/5527c7ecc9b943a8838ba6cba9a3b4d3.E_input', d5_relat_1)).
fof(t142_funct_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>(r2_hidden(X1,k2_relat_1(X2))<=>~(k10_relat_1(X2,k1_tarski(X1))=k1_xboole_0))), file('data/MPTP/proofs/5527c7ecc9b943a8838ba6cba9a3b4d3.E_input', t142_funct_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/5527c7ecc9b943a8838ba6cba9a3b4d3.E_input', t113_zfmisc_1)).
fof(t173_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(k10_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k2_relat_1(X2),X1))), file('data/MPTP/proofs/5527c7ecc9b943a8838ba6cba9a3b4d3.E_input', t173_relat_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/5527c7ecc9b943a8838ba6cba9a3b4d3.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/5527c7ecc9b943a8838ba6cba9a3b4d3.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/5527c7ecc9b943a8838ba6cba9a3b4d3.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/5527c7ecc9b943a8838ba6cba9a3b4d3.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/5527c7ecc9b943a8838ba6cba9a3b4d3.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/5527c7ecc9b943a8838ba6cba9a3b4d3.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/5527c7ecc9b943a8838ba6cba9a3b4d3.E_input', t75_enumset1)).
fof(t152_zfmisc_1, axiom, ![X1, X2]:~(r2_hidden(X1,k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/5527c7ecc9b943a8838ba6cba9a3b4d3.E_input', t152_zfmisc_1)).
fof(d3_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:(X6=k3_enumset1(X1,X2,X3,X4,X5)<=>![X7]:(r2_hidden(X7,X6)<=>~((~(X7=X1)&(~(X7=X2)&(~(X7=X3)&(~(X7=X4)&~(X7=X5)))))))), file('data/MPTP/proofs/5527c7ecc9b943a8838ba6cba9a3b4d3.E_input', d3_enumset1)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/5527c7ecc9b943a8838ba6cba9a3b4d3.E_input', t83_xboole_1)).
fof(l27_zfmisc_1, axiom, ![X1, X2]:(~(r2_hidden(X1,X2))=>r1_xboole_0(k1_tarski(X1),X2)), file('data/MPTP/proofs/5527c7ecc9b943a8838ba6cba9a3b4d3.E_input', l27_zfmisc_1)).
fof(c_0_18, plain, ![X8, X9]:(~r1_xboole_0(X8,X9)|r1_xboole_0(X9,X8)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
fof(c_0_19, plain, ![X10, X11]:r1_xboole_0(k4_xboole_0(X10,X11),X11), inference(variable_rename,[status(thm)],[t79_xboole_1])).
fof(c_0_20, plain, ![X12, X13, X14, X15, X17, X18, X19, X20, X22]:((((r2_hidden(k4_tarski(X15,esk1_4(X12,X13,X14,X15)),X12)|~r2_hidden(X15,X14)|X14!=k10_relat_1(X12,X13)|~v1_relat_1(X12))&(r2_hidden(esk1_4(X12,X13,X14,X15),X13)|~r2_hidden(X15,X14)|X14!=k10_relat_1(X12,X13)|~v1_relat_1(X12)))&(~r2_hidden(k4_tarski(X17,X18),X12)|~r2_hidden(X18,X13)|r2_hidden(X17,X14)|X14!=k10_relat_1(X12,X13)|~v1_relat_1(X12)))&((~r2_hidden(esk2_3(X12,X19,X20),X20)|(~r2_hidden(k4_tarski(esk2_3(X12,X19,X20),X22),X12)|~r2_hidden(X22,X19))|X20=k10_relat_1(X12,X19)|~v1_relat_1(X12))&((r2_hidden(k4_tarski(esk2_3(X12,X19,X20),esk3_3(X12,X19,X20)),X12)|r2_hidden(esk2_3(X12,X19,X20),X20)|X20=k10_relat_1(X12,X19)|~v1_relat_1(X12))&(r2_hidden(esk3_3(X12,X19,X20),X19)|r2_hidden(esk2_3(X12,X19,X20),X20)|X20=k10_relat_1(X12,X19)|~v1_relat_1(X12))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d14_relat_1])])])])])])).
fof(c_0_21, plain, ![X24, X25, X26, X28, X29, X30, X31, X33]:(((~r2_hidden(X26,X25)|r2_hidden(k4_tarski(esk4_3(X24,X25,X26),X26),X24)|X25!=k2_relat_1(X24))&(~r2_hidden(k4_tarski(X29,X28),X24)|r2_hidden(X28,X25)|X25!=k2_relat_1(X24)))&((~r2_hidden(esk5_2(X30,X31),X31)|~r2_hidden(k4_tarski(X33,esk5_2(X30,X31)),X30)|X31=k2_relat_1(X30))&(r2_hidden(esk5_2(X30,X31),X31)|r2_hidden(k4_tarski(esk6_2(X30,X31),esk5_2(X30,X31)),X30)|X31=k2_relat_1(X30)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
fof(c_0_22, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>(r2_hidden(X1,k2_relat_1(X2))<=>~(k10_relat_1(X2,k1_tarski(X1))=k1_xboole_0)))), inference(assume_negation,[status(cth)],[t142_funct_1])).
fof(c_0_23, plain, ![X37, X38]:((k2_zfmisc_1(X37,X38)!=k1_xboole_0|(X37=k1_xboole_0|X38=k1_xboole_0))&((X37!=k1_xboole_0|k2_zfmisc_1(X37,X38)=k1_xboole_0)&(X38!=k1_xboole_0|k2_zfmisc_1(X37,X38)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_24, plain, ![X39, X40]:((k10_relat_1(X40,X39)!=k1_xboole_0|r1_xboole_0(k2_relat_1(X40),X39)|~v1_relat_1(X40))&(~r1_xboole_0(k2_relat_1(X40),X39)|k10_relat_1(X40,X39)=k1_xboole_0|~v1_relat_1(X40))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t173_relat_1])])])).
cnf(c_0_25, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, plain, (r1_xboole_0(k4_xboole_0(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (r2_hidden(X1,X5)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X2,X4)|X5!=k10_relat_1(X3,X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (r2_hidden(k4_tarski(esk4_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_29, negated_conjecture, (v1_relat_1(esk8_0)&((~r2_hidden(esk7_0,k2_relat_1(esk8_0))|k10_relat_1(esk8_0,k1_tarski(esk7_0))=k1_xboole_0)&(r2_hidden(esk7_0,k2_relat_1(esk8_0))|k10_relat_1(esk8_0,k1_tarski(esk7_0))!=k1_xboole_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_22])])])])).
fof(c_0_30, plain, ![X41]:k2_tarski(X41,X41)=k1_tarski(X41), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_31, plain, ![X42, X43]:k1_enumset1(X42,X42,X43)=k2_tarski(X42,X43), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_32, plain, ![X44, X45, X46]:k2_enumset1(X44,X44,X45,X46)=k1_enumset1(X44,X45,X46), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_33, plain, ![X47, X48, X49, X50]:k3_enumset1(X47,X47,X48,X49,X50)=k2_enumset1(X47,X48,X49,X50), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_34, plain, ![X51, X52, X53, X54, X55]:k4_enumset1(X51,X51,X52,X53,X54,X55)=k3_enumset1(X51,X52,X53,X54,X55), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_35, plain, ![X56, X57, X58, X59, X60, X61]:k5_enumset1(X56,X56,X57,X58,X59,X60,X61)=k4_enumset1(X56,X57,X58,X59,X60,X61), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_36, plain, ![X62, X63, X64, X65, X66, X67, X68]:k6_enumset1(X62,X62,X63,X64,X65,X66,X67,X68)=k5_enumset1(X62,X63,X64,X65,X66,X67,X68), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_37, plain, ![X69, X70]:~r2_hidden(X69,k2_zfmisc_1(X69,X70)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_zfmisc_1])])).
cnf(c_0_38, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_39, plain, ![X71, X72, X73, X74, X75, X76, X77, X78, X79, X80, X81, X82, X83, X84]:(((~r2_hidden(X77,X76)|(X77=X71|(X77=X72|(X77=X73|(X77=X74|X77=X75))))|X76!=k3_enumset1(X71,X72,X73,X74,X75))&((X78!=X71|r2_hidden(X78,X76)|X76!=k3_enumset1(X71,X72,X73,X74,X75))&((X78!=X72|r2_hidden(X78,X76)|X76!=k3_enumset1(X71,X72,X73,X74,X75))&((X78!=X73|r2_hidden(X78,X76)|X76!=k3_enumset1(X71,X72,X73,X74,X75))&((X78!=X74|r2_hidden(X78,X76)|X76!=k3_enumset1(X71,X72,X73,X74,X75))&(X78!=X75|r2_hidden(X78,X76)|X76!=k3_enumset1(X71,X72,X73,X74,X75)))))))&(((esk9_6(X79,X80,X81,X82,X83,X84)!=X79|~r2_hidden(esk9_6(X79,X80,X81,X82,X83,X84),X84)|X84=k3_enumset1(X79,X80,X81,X82,X83))&((esk9_6(X79,X80,X81,X82,X83,X84)!=X80|~r2_hidden(esk9_6(X79,X80,X81,X82,X83,X84),X84)|X84=k3_enumset1(X79,X80,X81,X82,X83))&((esk9_6(X79,X80,X81,X82,X83,X84)!=X81|~r2_hidden(esk9_6(X79,X80,X81,X82,X83,X84),X84)|X84=k3_enumset1(X79,X80,X81,X82,X83))&((esk9_6(X79,X80,X81,X82,X83,X84)!=X82|~r2_hidden(esk9_6(X79,X80,X81,X82,X83,X84),X84)|X84=k3_enumset1(X79,X80,X81,X82,X83))&(esk9_6(X79,X80,X81,X82,X83,X84)!=X83|~r2_hidden(esk9_6(X79,X80,X81,X82,X83,X84),X84)|X84=k3_enumset1(X79,X80,X81,X82,X83))))))&(r2_hidden(esk9_6(X79,X80,X81,X82,X83,X84),X84)|(esk9_6(X79,X80,X81,X82,X83,X84)=X79|(esk9_6(X79,X80,X81,X82,X83,X84)=X80|(esk9_6(X79,X80,X81,X82,X83,X84)=X81|(esk9_6(X79,X80,X81,X82,X83,X84)=X82|esk9_6(X79,X80,X81,X82,X83,X84)=X83))))|X84=k3_enumset1(X79,X80,X81,X82,X83)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d3_enumset1])])])])])])])).
cnf(c_0_40, plain, (k10_relat_1(X1,X2)=k1_xboole_0|~r1_xboole_0(k2_relat_1(X1),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_41, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_42, plain, ![X86, X87]:((~r1_xboole_0(X86,X87)|k4_xboole_0(X86,X87)=X86)&(k4_xboole_0(X86,X87)!=X86|r1_xboole_0(X86,X87))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
fof(c_0_43, plain, ![X88, X89]:(r2_hidden(X88,X89)|r1_xboole_0(k1_tarski(X88),X89)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l27_zfmisc_1])])])).
cnf(c_0_44, plain, (r2_hidden(X1,k10_relat_1(X2,X3))|~r2_hidden(k4_tarski(X1,X4),X2)|~r2_hidden(X4,X3)|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_45, plain, (r2_hidden(k4_tarski(esk4_3(X1,k2_relat_1(X1),X2),X2),X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_28])).
cnf(c_0_46, negated_conjecture, (k10_relat_1(esk8_0,k1_tarski(esk7_0))=k1_xboole_0|~r2_hidden(esk7_0,k2_relat_1(esk8_0))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_47, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_48, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_49, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_50, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_51, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_52, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_53, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_54, plain, (~r2_hidden(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_55, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_56, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k3_enumset1(X4,X2,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_57, plain, (k10_relat_1(X1,k4_xboole_0(X2,k2_relat_1(X1)))=k1_xboole_0|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_58, plain, (k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_59, plain, (r2_hidden(X1,X2)|r1_xboole_0(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_60, plain, (r2_hidden(esk4_3(X1,k2_relat_1(X1),X2),k10_relat_1(X1,X3))|~r2_hidden(X2,k2_relat_1(X1))|~r2_hidden(X2,X3)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_61, negated_conjecture, (k10_relat_1(esk8_0,k6_enumset1(esk7_0,esk7_0,esk7_0,esk7_0,esk7_0,esk7_0,esk7_0,esk7_0))=k1_xboole_0|~r2_hidden(esk7_0,k2_relat_1(esk8_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_47]), c_0_48]), c_0_49]), c_0_50]), c_0_51]), c_0_52]), c_0_53])).
cnf(c_0_62, negated_conjecture, (v1_relat_1(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_63, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_64, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k6_enumset1(X4,X4,X4,X4,X2,X5,X6,X7)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_51]), c_0_52]), c_0_53])).
cnf(c_0_65, negated_conjecture, (r2_hidden(esk7_0,k2_relat_1(esk8_0))|k10_relat_1(esk8_0,k1_tarski(esk7_0))!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_66, plain, (k10_relat_1(X1,X2)=k1_xboole_0|~v1_relat_1(X1)|~r1_xboole_0(X2,k2_relat_1(X1))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_67, plain, (r2_hidden(X1,X2)|r1_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_47]), c_0_48]), c_0_49]), c_0_50]), c_0_51]), c_0_52]), c_0_53])).
cnf(c_0_68, negated_conjecture, (~r2_hidden(X1,k6_enumset1(esk7_0,esk7_0,esk7_0,esk7_0,esk7_0,esk7_0,esk7_0,esk7_0))|~r2_hidden(esk7_0,k2_relat_1(esk8_0))|~r2_hidden(X1,k2_relat_1(esk8_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62])]), c_0_63])).
cnf(c_0_69, plain, (r2_hidden(X1,k6_enumset1(X2,X2,X2,X2,X1,X3,X4,X5))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_64])])).
cnf(c_0_70, negated_conjecture, (r2_hidden(esk7_0,k2_relat_1(esk8_0))|k10_relat_1(esk8_0,k6_enumset1(esk7_0,esk7_0,esk7_0,esk7_0,esk7_0,esk7_0,esk7_0,esk7_0))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_47]), c_0_48]), c_0_49]), c_0_50]), c_0_51]), c_0_52]), c_0_53])).
cnf(c_0_71, plain, (k10_relat_1(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2))=k1_xboole_0|r2_hidden(X2,k2_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_72, negated_conjecture, (~r2_hidden(esk7_0,k2_relat_1(esk8_0))), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_73, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_62])]), c_0_72]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 74
# Proof object clause steps            : 37
# Proof object formula steps           : 37
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 35
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 43
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 36
# Processed clauses                    : 424
# ...of these trivial                  : 0
# ...subsumed                          : 246
# ...remaining for further processing  : 178
# Other redundant clauses eliminated   : 24
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 19
# Generated clauses                    : 1325
# ...of the previous two non-trivial   : 1151
# Contextual simplify-reflections      : 0
# Paramodulations                      : 1306
# Factorizations                       : 0
# Equation resolutions                 : 24
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 107
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 87
# Current number of unprocessed clauses: 755
# ...number of literals in the above   : 4105
# Current number of archived formulas  : 0
# Current number of archived clauses   : 65
# Clause-clause subsumption calls (NU) : 818
# Rec. Clause-clause subsumption calls : 748
# Non-unit clause-clause subsumptions  : 38
# Unit Clause-clause subsumption calls : 17
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 31
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 29269

# -------------------------------------------------
# User time                : 0.028 s
# System time              : 0.000 s
# Total time               : 0.028 s
# Maximum resident set size: 28664 pages
