# No SInE strategy applied
# Auto-Mode selected heuristic H_____047_B31_F1_PI_AE_R4_CS_SP_S4c
# and selection function SelectCQPrecWNTNp.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t50_orders_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>~(r2_hidden(X2,k2_orders_2(X1,k6_domain_1(u1_struct_0(X1),X2)))))), file('data/MPTP/proofs/556356847f55440c88bfdcb193accc26.E_input', t50_orders_2)).
fof(t35_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v6_orders_2(k6_domain_1(u1_struct_0(X1),X2),X1)&m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/556356847f55440c88bfdcb193accc26.E_input', t35_orders_2)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/556356847f55440c88bfdcb193accc26.E_input', redefinition_k6_domain_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/556356847f55440c88bfdcb193accc26.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/556356847f55440c88bfdcb193accc26.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/556356847f55440c88bfdcb193accc26.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/556356847f55440c88bfdcb193accc26.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/556356847f55440c88bfdcb193accc26.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/556356847f55440c88bfdcb193accc26.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/556356847f55440c88bfdcb193accc26.E_input', t75_enumset1)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/556356847f55440c88bfdcb193accc26.E_input', t5_subset)).
fof(dt_k2_orders_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k2_orders_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/556356847f55440c88bfdcb193accc26.E_input', dt_k2_orders_2)).
fof(d6_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:(X9=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)<=>![X10]:(r2_hidden(X10,X9)<=>~((~(X10=X1)&(~(X10=X2)&(~(X10=X3)&(~(X10=X4)&(~(X10=X5)&(~(X10=X6)&(~(X10=X7)&~(X10=X8))))))))))), file('data/MPTP/proofs/556356847f55440c88bfdcb193accc26.E_input', d6_enumset1)).
fof(t49_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>~((r2_hidden(X2,X3)&r2_hidden(X2,k2_orders_2(X1,X3))))))), file('data/MPTP/proofs/556356847f55440c88bfdcb193accc26.E_input', t49_orders_2)).
fof(c_0_14, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>~(r2_hidden(X2,k2_orders_2(X1,k6_domain_1(u1_struct_0(X1),X2))))))), inference(assume_negation,[status(cth)],[t50_orders_2])).
fof(c_0_15, plain, ![X13, X14]:((v6_orders_2(k6_domain_1(u1_struct_0(X13),X14),X13)|~m1_subset_1(X14,u1_struct_0(X13))|(v2_struct_0(X13)|(~v3_orders_2(X13)|~l1_orders_2(X13))))&(m1_subset_1(k6_domain_1(u1_struct_0(X13),X14),k1_zfmisc_1(u1_struct_0(X13)))|~m1_subset_1(X14,u1_struct_0(X13))|(v2_struct_0(X13)|(~v3_orders_2(X13)|~l1_orders_2(X13))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t35_orders_2])])])])])).
fof(c_0_16, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&r2_hidden(esk2_0,k2_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])).
fof(c_0_17, plain, ![X53, X54]:(v1_xboole_0(X53)|~m1_subset_1(X54,X53)|k6_domain_1(X53,X54)=k1_tarski(X54)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_18, plain, ![X55]:k2_tarski(X55,X55)=k1_tarski(X55), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_19, plain, ![X72, X73]:k1_enumset1(X72,X72,X73)=k2_tarski(X72,X73), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_20, plain, ![X24, X25, X26]:k2_enumset1(X24,X24,X25,X26)=k1_enumset1(X24,X25,X26), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_21, plain, ![X68, X69, X70, X71]:k3_enumset1(X68,X68,X69,X70,X71)=k2_enumset1(X68,X69,X70,X71), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_22, plain, ![X27, X28, X29, X30, X31]:k4_enumset1(X27,X27,X28,X29,X30,X31)=k3_enumset1(X27,X28,X29,X30,X31), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_23, plain, ![X56, X57, X58, X59, X60, X61]:k5_enumset1(X56,X56,X57,X58,X59,X60,X61)=k4_enumset1(X56,X57,X58,X59,X60,X61), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_24, plain, ![X15, X16, X17, X18, X19, X20, X21]:k6_enumset1(X15,X15,X16,X17,X18,X19,X20,X21)=k5_enumset1(X15,X16,X17,X18,X19,X20,X21), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_25, plain, ![X65, X66, X67]:(~r2_hidden(X65,X66)|(~m1_subset_1(X66,k1_zfmisc_1(X67))|~v1_xboole_0(X67))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_26, plain, ![X22, X23]:(v2_struct_0(X22)|(~v3_orders_2(X22)|(~v4_orders_2(X22)|(~v5_orders_2(X22)|~l1_orders_2(X22))))|~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(X22)))|m1_subset_1(k2_orders_2(X22,X23),k1_zfmisc_1(u1_struct_0(X22)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_orders_2])])])).
cnf(c_0_27, plain, (m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_28, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_29, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_30, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_31, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_32, plain, ![X32, X33, X34, X35, X36, X37, X38, X39, X40, X41, X42, X43, X44, X45, X46, X47, X48, X49, X50, X51]:(((~r2_hidden(X41,X40)|(X41=X32|(X41=X33|(X41=X34|(X41=X35|(X41=X36|(X41=X37|(X41=X38|X41=X39)))))))|X40!=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((X42!=X32|r2_hidden(X42,X40)|X40!=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((X42!=X33|r2_hidden(X42,X40)|X40!=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((X42!=X34|r2_hidden(X42,X40)|X40!=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((X42!=X35|r2_hidden(X42,X40)|X40!=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((X42!=X36|r2_hidden(X42,X40)|X40!=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((X42!=X37|r2_hidden(X42,X40)|X40!=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((X42!=X38|r2_hidden(X42,X40)|X40!=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&(X42!=X39|r2_hidden(X42,X40)|X40!=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))))))))))&(((esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)!=X43|~r2_hidden(esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51),X51)|X51=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&((esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)!=X44|~r2_hidden(esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51),X51)|X51=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&((esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)!=X45|~r2_hidden(esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51),X51)|X51=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&((esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)!=X46|~r2_hidden(esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51),X51)|X51=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&((esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)!=X47|~r2_hidden(esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51),X51)|X51=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&((esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)!=X48|~r2_hidden(esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51),X51)|X51=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&((esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)!=X49|~r2_hidden(esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51),X51)|X51=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&(esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)!=X50|~r2_hidden(esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51),X51)|X51=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50)))))))))&(r2_hidden(esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51),X51)|(esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)=X43|(esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)=X44|(esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)=X45|(esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)=X46|(esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)=X47|(esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)=X48|(esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)=X49|esk3_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)=X50)))))))|X51=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_enumset1])])])])])])])).
cnf(c_0_33, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_34, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_35, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_36, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_37, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_38, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_39, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_40, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_41, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_42, negated_conjecture, (r2_hidden(esk2_0,k2_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_43, plain, (v2_struct_0(X1)|m1_subset_1(k2_orders_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_44, negated_conjecture, (m1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_45, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_46, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_47, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k6_enumset1(X4,X5,X6,X7,X8,X9,X10,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_48, plain, (k6_domain_1(X1,X2)=k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40])).
cnf(c_0_49, negated_conjecture, (~v1_xboole_0(X1)|~m1_subset_1(k2_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0)),k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_50, negated_conjecture, (m1_subset_1(k2_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_29]), c_0_45]), c_0_46]), c_0_30])]), c_0_31])).
fof(c_0_51, plain, ![X62, X63, X64]:(v2_struct_0(X62)|(~v3_orders_2(X62)|(~v4_orders_2(X62)|(~v5_orders_2(X62)|~l1_orders_2(X62))))|(~m1_subset_1(X63,u1_struct_0(X62))|(~m1_subset_1(X64,k1_zfmisc_1(u1_struct_0(X62)))|(~r2_hidden(X63,X64)|~r2_hidden(X63,k2_orders_2(X62,X64)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t49_orders_2])])])])).
cnf(c_0_52, plain, (r2_hidden(X1,X2)|X2!=k6_enumset1(X3,X4,X5,X6,X7,X8,X9,X1)), inference(er,[status(thm)],[c_0_47])).
cnf(c_0_53, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_48, c_0_28])).
cnf(c_0_54, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_55, plain, (v2_struct_0(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,X3)|~r2_hidden(X2,k2_orders_2(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_56, plain, (r2_hidden(X1,k6_enumset1(X2,X3,X4,X5,X6,X7,X8,X1))), inference(er,[status(thm)],[c_0_52])).
cnf(c_0_57, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)), inference(sr,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_58, plain, (v2_struct_0(X1)|~r2_hidden(X2,k2_orders_2(X1,k6_enumset1(X3,X4,X5,X6,X7,X8,X9,X2)))|~m1_subset_1(k6_enumset1(X3,X4,X5,X6,X7,X8,X9,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_59, negated_conjecture, (r2_hidden(esk2_0,k2_orders_2(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))), inference(rw,[status(thm)],[c_0_42, c_0_57])).
cnf(c_0_60, negated_conjecture, (m1_subset_1(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(rw,[status(thm)],[c_0_44, c_0_57])).
cnf(c_0_61, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_60]), c_0_28]), c_0_29]), c_0_45]), c_0_46]), c_0_30])]), c_0_31]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 62
# Proof object clause steps            : 33
# Proof object formula steps           : 29
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 29
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 38
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 31
# Processed clauses                    : 79
# ...of these trivial                  : 0
# ...subsumed                          : 8
# ...remaining for further processing  : 71
# Other redundant clauses eliminated   : 17
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 6
# Generated clauses                    : 192
# ...of the previous two non-trivial   : 178
# Contextual simplify-reflections      : 0
# Paramodulations                      : 165
# Factorizations                       : 0
# Equation resolutions                 : 26
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 56
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 37
# Current number of unprocessed clauses: 127
# ...number of literals in the above   : 1564
# Current number of archived formulas  : 0
# Current number of archived clauses   : 14
# Clause-clause subsumption calls (NU) : 439
# Rec. Clause-clause subsumption calls : 120
# Non-unit clause-clause subsumptions  : 7
# Unit Clause-clause subsumption calls : 14
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 32
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7309

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.003 s
# Total time               : 0.016 s
# Maximum resident set size: 28700 pages
