# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/557684b336594631b686695a3ba08bc4.E_input', t4_xboole_0)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/557684b336594631b686695a3ba08bc4.E_input', t12_setfam_1)).
fof(t166_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k10_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k2_relat_1(X3))&(r2_hidden(k4_tarski(X1,X4),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/557684b336594631b686695a3ba08bc4.E_input', t166_relat_1)).
fof(t168_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k10_relat_1(X2,X1)=k10_relat_1(X2,k3_xboole_0(k2_relat_1(X2),X1))), file('data/MPTP/proofs/557684b336594631b686695a3ba08bc4.E_input', t168_relat_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/557684b336594631b686695a3ba08bc4.E_input', t3_xboole_0)).
fof(t173_relat_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>(k10_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k2_relat_1(X2),X1))), file('data/MPTP/proofs/557684b336594631b686695a3ba08bc4.E_input', t173_relat_1)).
fof(l24_zfmisc_1, axiom, ![X1, X2]:~((r1_xboole_0(k1_tarski(X1),X2)&r2_hidden(X1,X2))), file('data/MPTP/proofs/557684b336594631b686695a3ba08bc4.E_input', l24_zfmisc_1)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/557684b336594631b686695a3ba08bc4.E_input', t65_xboole_1)).
fof(d14_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(X3=k10_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(k4_tarski(X4,X5),X1)&r2_hidden(X5,X2))))), file('data/MPTP/proofs/557684b336594631b686695a3ba08bc4.E_input', d14_relat_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/557684b336594631b686695a3ba08bc4.E_input', d5_relat_1)).
fof(t60_relat_1, axiom, (k1_relat_1(k1_xboole_0)=k1_xboole_0&k2_relat_1(k1_xboole_0)=k1_xboole_0), file('data/MPTP/proofs/557684b336594631b686695a3ba08bc4.E_input', t60_relat_1)).
fof(c_0_11, plain, ![X47, X48, X50, X51, X52]:((r1_xboole_0(X47,X48)|r2_hidden(esk11_2(X47,X48),k3_xboole_0(X47,X48)))&(~r2_hidden(X52,k3_xboole_0(X50,X51))|~r1_xboole_0(X50,X51))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_12, plain, ![X53, X54]:k1_setfam_1(k2_tarski(X53,X54))=k3_xboole_0(X53,X54), inference(variable_rename,[status(thm)],[t12_setfam_1])).
cnf(c_0_13, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_14, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_15, plain, ![X23, X24, X25, X27]:(((r2_hidden(esk6_3(X23,X24,X25),k2_relat_1(X25))|~r2_hidden(X23,k10_relat_1(X25,X24))|~v1_relat_1(X25))&((r2_hidden(k4_tarski(X23,esk6_3(X23,X24,X25)),X25)|~r2_hidden(X23,k10_relat_1(X25,X24))|~v1_relat_1(X25))&(r2_hidden(esk6_3(X23,X24,X25),X24)|~r2_hidden(X23,k10_relat_1(X25,X24))|~v1_relat_1(X25))))&(~r2_hidden(X27,k2_relat_1(X25))|(~r2_hidden(k4_tarski(X23,X27),X25)|~r2_hidden(X27,X24))|r2_hidden(X23,k10_relat_1(X25,X24))|~v1_relat_1(X25))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t166_relat_1])])])])])).
fof(c_0_16, plain, ![X28, X29]:(~v1_relat_1(X29)|k10_relat_1(X29,X28)=k10_relat_1(X29,k3_xboole_0(k2_relat_1(X29),X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t168_relat_1])])).
cnf(c_0_17, plain, (~r1_xboole_0(X2,X3)|~r2_hidden(X1,k1_setfam_1(k2_tarski(X2,X3)))), inference(rw,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_18, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, plain, (k10_relat_1(X1,X2)=k10_relat_1(X1,k3_xboole_0(k2_relat_1(X1),X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, (~r2_hidden(X1,k10_relat_1(X2,k1_setfam_1(k2_tarski(X3,X4))))|~r1_xboole_0(X3,X4)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_21, plain, (k10_relat_1(X1,X2)=k10_relat_1(X1,k1_setfam_1(k2_tarski(k2_relat_1(X1),X2)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_19, c_0_14])).
fof(c_0_22, plain, ![X30, X31, X33, X34, X35]:(((r2_hidden(esk7_2(X30,X31),X30)|r1_xboole_0(X30,X31))&(r2_hidden(esk7_2(X30,X31),X31)|r1_xboole_0(X30,X31)))&(~r2_hidden(X35,X33)|~r2_hidden(X35,X34)|~r1_xboole_0(X33,X34))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_23, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>(k10_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k2_relat_1(X2),X1)))), inference(assume_negation,[status(cth)],[t173_relat_1])).
cnf(c_0_24, plain, (~r2_hidden(X1,k10_relat_1(X2,X3))|~r1_xboole_0(k2_relat_1(X2),X3)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_25, plain, (r2_hidden(esk7_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_26, negated_conjecture, (v1_relat_1(esk2_0)&((k10_relat_1(esk2_0,esk1_0)!=k1_xboole_0|~r1_xboole_0(k2_relat_1(esk2_0),esk1_0))&(k10_relat_1(esk2_0,esk1_0)=k1_xboole_0|r1_xboole_0(k2_relat_1(esk2_0),esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])).
fof(c_0_27, plain, ![X9, X10]:(~r1_xboole_0(k1_tarski(X9),X10)|~r2_hidden(X9,X10)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l24_zfmisc_1])])).
fof(c_0_28, plain, ![X8]:r1_xboole_0(X8,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
fof(c_0_29, plain, ![X11, X12, X13, X14, X16, X17, X18, X19, X21]:((((r2_hidden(k4_tarski(X14,esk3_4(X11,X12,X13,X14)),X11)|~r2_hidden(X14,X13)|X13!=k10_relat_1(X11,X12)|~v1_relat_1(X11))&(r2_hidden(esk3_4(X11,X12,X13,X14),X12)|~r2_hidden(X14,X13)|X13!=k10_relat_1(X11,X12)|~v1_relat_1(X11)))&(~r2_hidden(k4_tarski(X16,X17),X11)|~r2_hidden(X17,X12)|r2_hidden(X16,X13)|X13!=k10_relat_1(X11,X12)|~v1_relat_1(X11)))&((~r2_hidden(esk4_3(X11,X18,X19),X19)|(~r2_hidden(k4_tarski(esk4_3(X11,X18,X19),X21),X11)|~r2_hidden(X21,X18))|X19=k10_relat_1(X11,X18)|~v1_relat_1(X11))&((r2_hidden(k4_tarski(esk4_3(X11,X18,X19),esk5_3(X11,X18,X19)),X11)|r2_hidden(esk4_3(X11,X18,X19),X19)|X19=k10_relat_1(X11,X18)|~v1_relat_1(X11))&(r2_hidden(esk5_3(X11,X18,X19),X18)|r2_hidden(esk4_3(X11,X18,X19),X19)|X19=k10_relat_1(X11,X18)|~v1_relat_1(X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d14_relat_1])])])])])])).
fof(c_0_30, plain, ![X36, X37, X38, X40, X41, X42, X43, X45]:(((~r2_hidden(X38,X37)|r2_hidden(k4_tarski(esk8_3(X36,X37,X38),X38),X36)|X37!=k2_relat_1(X36))&(~r2_hidden(k4_tarski(X41,X40),X36)|r2_hidden(X40,X37)|X37!=k2_relat_1(X36)))&((~r2_hidden(esk9_2(X42,X43),X43)|~r2_hidden(k4_tarski(X45,esk9_2(X42,X43)),X42)|X43=k2_relat_1(X42))&(r2_hidden(esk9_2(X42,X43),X43)|r2_hidden(k4_tarski(esk10_2(X42,X43),esk9_2(X42,X43)),X42)|X43=k2_relat_1(X42)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_31, plain, (r1_xboole_0(X1,k10_relat_1(X2,X3))|~r1_xboole_0(k2_relat_1(X2),X3)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_32, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)=k1_xboole_0|r1_xboole_0(k2_relat_1(esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (~r1_xboole_0(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36, plain, (r2_hidden(X1,X5)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X2,X4)|X5!=k10_relat_1(X3,X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, plain, (r2_hidden(k4_tarski(esk8_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)=k1_xboole_0|r1_xboole_0(X1,k10_relat_1(esk2_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])])).
cnf(c_0_39, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40, plain, (r2_hidden(esk9_2(X1,X2),X2)|r2_hidden(k4_tarski(esk10_2(X1,X2),esk9_2(X1,X2)),X1)|X2=k2_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_41, plain, (k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
cnf(c_0_42, plain, (r2_hidden(X1,k10_relat_1(X2,X3))|~r2_hidden(k4_tarski(X1,X4),X2)|~r2_hidden(X4,X3)|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_36])).
cnf(c_0_43, plain, (r2_hidden(k4_tarski(esk8_3(X1,k2_relat_1(X1),X2),X2),X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_44, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)=k1_xboole_0|~r2_hidden(X1,k10_relat_1(esk2_0,esk1_0))), inference(spm,[status(thm)],[c_0_34, c_0_38])).
cnf(c_0_45, plain, (X1=k1_xboole_0|r2_hidden(esk9_2(k1_xboole_0,X1),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])).
cnf(c_0_46, plain, (r2_hidden(esk8_3(X1,k2_relat_1(X1),X2),k10_relat_1(X1,X3))|~r2_hidden(X2,k2_relat_1(X1))|~r2_hidden(X2,X3)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_47, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_48, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_49, plain, (r2_hidden(esk7_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_50, negated_conjecture, (~r2_hidden(X1,k2_relat_1(esk2_0))|~r2_hidden(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_33])]), c_0_39])).
cnf(c_0_51, plain, (r1_xboole_0(X1,X2)|~r2_hidden(esk7_2(X1,X2),X3)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_52, negated_conjecture, (r1_xboole_0(X1,k2_relat_1(esk2_0))|~r2_hidden(esk7_2(X1,k2_relat_1(esk2_0)),esk1_0)), inference(spm,[status(thm)],[c_0_50, c_0_25])).
cnf(c_0_53, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)!=k1_xboole_0|~r1_xboole_0(k2_relat_1(esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_54, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_51, c_0_25])).
cnf(c_0_55, negated_conjecture, (r1_xboole_0(esk1_0,k2_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_52, c_0_49])).
cnf(c_0_56, negated_conjecture, (~r1_xboole_0(k2_relat_1(esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_47])])).
cnf(c_0_57, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 58
# Proof object clause steps            : 36
# Proof object formula steps           : 22
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 13
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 27
# Processed clauses                    : 3774
# ...of these trivial                  : 0
# ...subsumed                          : 3271
# ...remaining for further processing  : 503
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 35
# Generated clauses                    : 21101
# ...of the previous two non-trivial   : 20462
# Contextual simplify-reflections      : 0
# Paramodulations                      : 21094
# Factorizations                       : 2
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 433
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 423
# Current number of unprocessed clauses: 16592
# ...number of literals in the above   : 70667
# Current number of archived formulas  : 0
# Current number of archived clauses   : 66
# Clause-clause subsumption calls (NU) : 76754
# Rec. Clause-clause subsumption calls : 48739
# Non-unit clause-clause subsumptions  : 2774
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 363859

# -------------------------------------------------
# User time                : 0.204 s
# System time              : 0.000 s
# Total time               : 0.204 s
# Maximum resident set size: 28592 pages
