# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/55f89e7b2b5443e3a9cc4ab353dbbafb.E_input', t1_boole)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/55f89e7b2b5443e3a9cc4ab353dbbafb.E_input', commutativity_k2_xboole_0)).
fof(t22_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/55f89e7b2b5443e3a9cc4ab353dbbafb.E_input', t22_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/55f89e7b2b5443e3a9cc4ab353dbbafb.E_input', t28_xboole_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/55f89e7b2b5443e3a9cc4ab353dbbafb.E_input', t36_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/55f89e7b2b5443e3a9cc4ab353dbbafb.E_input', commutativity_k3_xboole_0)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/55f89e7b2b5443e3a9cc4ab353dbbafb.E_input', t17_xboole_1)).
fof(t16_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3)), file('data/MPTP/proofs/55f89e7b2b5443e3a9cc4ab353dbbafb.E_input', t16_xboole_1)).
fof(t21_xboole_1, axiom, ![X1, X2]:k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/55f89e7b2b5443e3a9cc4ab353dbbafb.E_input', t21_xboole_1)).
fof(t40_xboole_1, axiom, ![X1, X2]:k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/55f89e7b2b5443e3a9cc4ab353dbbafb.E_input', t40_xboole_1)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/55f89e7b2b5443e3a9cc4ab353dbbafb.E_input', t41_xboole_1)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/55f89e7b2b5443e3a9cc4ab353dbbafb.E_input', t39_xboole_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/55f89e7b2b5443e3a9cc4ab353dbbafb.E_input', l32_xboole_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/55f89e7b2b5443e3a9cc4ab353dbbafb.E_input', d5_xboole_0)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/55f89e7b2b5443e3a9cc4ab353dbbafb.E_input', idempotence_k3_xboole_0)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/55f89e7b2b5443e3a9cc4ab353dbbafb.E_input', t3_xboole_0)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/55f89e7b2b5443e3a9cc4ab353dbbafb.E_input', d7_xboole_0)).
fof(t47_xboole_1, conjecture, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/55f89e7b2b5443e3a9cc4ab353dbbafb.E_input', t47_xboole_1)).
fof(c_0_18, plain, ![X5]:k2_xboole_0(X5,k1_xboole_0)=X5, inference(variable_rename,[status(thm)],[t1_boole])).
fof(c_0_19, plain, ![X6, X7]:k2_xboole_0(X6,X7)=k2_xboole_0(X7,X6), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_20, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_21, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_22, plain, ![X29, X30]:k2_xboole_0(X29,k3_xboole_0(X29,X30))=X29, inference(variable_rename,[status(thm)],[t22_xboole_1])).
fof(c_0_23, plain, ![X10, X11]:(~r1_tarski(X10,X11)|k3_xboole_0(X10,X11)=X10), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_24, plain, ![X12, X13]:r1_tarski(k4_xboole_0(X12,X13),X12), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_25, plain, ![X14, X15]:k3_xboole_0(X14,X15)=k3_xboole_0(X15,X14), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_26, plain, ![X16, X17]:r1_tarski(k3_xboole_0(X16,X17),X16), inference(variable_rename,[status(thm)],[t17_xboole_1])).
cnf(c_0_27, plain, (k2_xboole_0(k1_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_28, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_33, plain, ![X18, X19, X20]:k3_xboole_0(k3_xboole_0(X18,X19),X20)=k3_xboole_0(X18,k3_xboole_0(X19,X20)), inference(variable_rename,[status(thm)],[t16_xboole_1])).
fof(c_0_34, plain, ![X8, X9]:k3_xboole_0(X8,k2_xboole_0(X8,X9))=X8, inference(variable_rename,[status(thm)],[t21_xboole_1])).
fof(c_0_35, plain, ![X21, X22]:k4_xboole_0(k2_xboole_0(X21,X22),X22)=k4_xboole_0(X21,X22), inference(variable_rename,[status(thm)],[t40_xboole_1])).
cnf(c_0_36, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_37, plain, (k3_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])).
cnf(c_0_38, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_32, c_0_31])).
cnf(c_0_39, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_40, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_41, plain, ![X24, X25, X26]:k4_xboole_0(k4_xboole_0(X24,X25),X26)=k4_xboole_0(X24,k2_xboole_0(X25,X26)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
cnf(c_0_42, plain, (k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_43, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_36, c_0_37])).
fof(c_0_44, plain, ![X27, X28]:k2_xboole_0(X27,k4_xboole_0(X28,X27))=k2_xboole_0(X27,X28), inference(variable_rename,[status(thm)],[t39_xboole_1])).
cnf(c_0_45, plain, (r1_tarski(k3_xboole_0(X1,k3_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_46, plain, (k3_xboole_0(X1,k2_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_40, c_0_21])).
fof(c_0_47, plain, ![X40, X41]:((k4_xboole_0(X40,X41)!=k1_xboole_0|r1_tarski(X40,X41))&(~r1_tarski(X40,X41)|k4_xboole_0(X40,X41)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_48, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_49, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_27]), c_0_43])).
cnf(c_0_50, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_51, plain, (r1_tarski(k3_xboole_0(X1,X2),k2_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
fof(c_0_52, plain, ![X31, X32, X33, X34, X35, X36, X37, X38]:((((r2_hidden(X34,X31)|~r2_hidden(X34,X33)|X33!=k4_xboole_0(X31,X32))&(~r2_hidden(X34,X32)|~r2_hidden(X34,X33)|X33!=k4_xboole_0(X31,X32)))&(~r2_hidden(X35,X31)|r2_hidden(X35,X32)|r2_hidden(X35,X33)|X33!=k4_xboole_0(X31,X32)))&((~r2_hidden(esk1_3(X36,X37,X38),X38)|(~r2_hidden(esk1_3(X36,X37,X38),X36)|r2_hidden(esk1_3(X36,X37,X38),X37))|X38=k4_xboole_0(X36,X37))&((r2_hidden(esk1_3(X36,X37,X38),X36)|r2_hidden(esk1_3(X36,X37,X38),X38)|X38=k4_xboole_0(X36,X37))&(~r2_hidden(esk1_3(X36,X37,X38),X37)|r2_hidden(esk1_3(X36,X37,X38),X38)|X38=k4_xboole_0(X36,X37))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_53, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_54, plain, (k4_xboole_0(X1,k2_xboole_0(X2,X1))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])).
cnf(c_0_55, plain, (r1_tarski(X1,k2_xboole_0(X2,k2_xboole_0(X3,X1)))), inference(spm,[status(thm)],[c_0_51, c_0_46])).
cnf(c_0_56, plain, (k2_xboole_0(X1,k4_xboole_0(X1,X2))=X1), inference(spm,[status(thm)],[c_0_28, c_0_37])).
fof(c_0_57, plain, ![X23]:k3_xboole_0(X23,X23)=X23, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_58, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
fof(c_0_59, plain, ![X42, X43, X45, X46, X47]:(((r2_hidden(esk2_2(X42,X43),X42)|r1_xboole_0(X42,X43))&(r2_hidden(esk2_2(X42,X43),X43)|r1_xboole_0(X42,X43)))&(~r2_hidden(X47,X45)|~r2_hidden(X47,X46)|~r1_xboole_0(X45,X46))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_60, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|k4_xboole_0(X1,k2_xboole_0(X2,X3))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_53, c_0_48])).
cnf(c_0_61, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_54, c_0_28])).
cnf(c_0_62, plain, (r1_tarski(k4_xboole_0(X1,X2),k2_xboole_0(X3,X1))), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_63, plain, (k2_xboole_0(X1,k3_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_28, c_0_31])).
cnf(c_0_64, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_65, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_58])).
cnf(c_0_66, plain, (r2_hidden(esk2_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_67, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(k2_xboole_0(X1,X2),X3),X1),X2)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_68, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_69, plain, (k3_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_39, c_0_64])).
fof(c_0_70, plain, ![X48, X49]:((~r1_xboole_0(X48,X49)|k3_xboole_0(X48,X49)=k1_xboole_0)&(k3_xboole_0(X48,X49)!=k1_xboole_0|r1_xboole_0(X48,X49))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_71, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r2_hidden(esk2_2(X1,k4_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_72, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_73, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(k2_xboole_0(X1,X2),X3),X2),X1)), inference(spm,[status(thm)],[c_0_67, c_0_21])).
cnf(c_0_74, plain, (k3_xboole_0(X1,k4_xboole_0(k3_xboole_0(X2,X1),X3))=k4_xboole_0(k3_xboole_0(X2,X1),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_68]), c_0_31])).
cnf(c_0_75, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_69, c_0_31])).
cnf(c_0_76, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_77, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_78, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(k2_xboole_0(X1,X2),X3),k4_xboole_0(X2,X1)),X1)), inference(spm,[status(thm)],[c_0_73, c_0_50])).
cnf(c_0_79, plain, (k3_xboole_0(k2_xboole_0(X1,X2),k4_xboole_0(X2,X3))=k4_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_74, c_0_46])).
cnf(c_0_80, plain, (k3_xboole_0(X1,k3_xboole_0(X2,k3_xboole_0(X1,X3)))=k3_xboole_0(X2,k3_xboole_0(X1,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_75]), c_0_39]), c_0_39])).
cnf(c_0_81, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_82, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(X2,k4_xboole_0(X1,X3))),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_48])).
cnf(c_0_83, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(X2,X3)),k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_30, c_0_48])).
cnf(c_0_84, plain, (k3_xboole_0(X1,k3_xboole_0(X2,k4_xboole_0(X1,X3)))=k3_xboole_0(X2,k4_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_80, c_0_37])).
cnf(c_0_85, plain, (k3_xboole_0(X1,k3_xboole_0(k4_xboole_0(X2,X1),X3))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_81]), c_0_36])).
cnf(c_0_86, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_82, c_0_27])).
cnf(c_0_87, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X2)=k1_xboole_0), inference(spm,[status(thm)],[c_0_54, c_0_63])).
cnf(c_0_88, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(X2,X3)),k4_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_83, c_0_21])).
cnf(c_0_89, plain, (r1_tarski(k2_xboole_0(X1,X2),X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_53, c_0_42])).
cnf(c_0_90, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X3))=k1_xboole_0), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_91, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X1)))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_86]), c_0_31])).
cnf(c_0_92, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,k2_xboole_0(X2,X3)),X2),X3)), inference(spm,[status(thm)],[c_0_60, c_0_87])).
cnf(c_0_93, plain, (r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X3,X2)))), inference(spm,[status(thm)],[c_0_88, c_0_63])).
cnf(c_0_94, plain, (k2_xboole_0(X1,X2)=X2|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_89]), c_0_31]), c_0_46])).
cnf(c_0_95, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,X1)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_96, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,k2_xboole_0(X2,X3)),X2),k4_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_92, c_0_50])).
cnf(c_0_97, plain, (k3_xboole_0(X1,k2_xboole_0(X2,k2_xboole_0(X3,X1)))=X1), inference(spm,[status(thm)],[c_0_29, c_0_55])).
cnf(c_0_98, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X3,X2)))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_29, c_0_93])).
cnf(c_0_99, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_56])).
cnf(c_0_100, plain, (r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_88, c_0_28])).
cnf(c_0_101, plain, (r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(k2_xboole_0(X3,X1),X2))), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_102, plain, (k4_xboole_0(X1,k3_xboole_0(X2,k4_xboole_0(X3,X1)))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_37])).
cnf(c_0_103, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X2,X3)))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_29, c_0_100])).
cnf(c_0_104, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,X2),X3),k4_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_101, c_0_63])).
cnf(c_0_105, plain, (k4_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X3,X1))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_106, plain, (k4_xboole_0(k2_xboole_0(X1,X2),k4_xboole_0(X2,X1))=k4_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_42, c_0_50])).
cnf(c_0_107, plain, (r1_tarski(k3_xboole_0(X1,X2),k4_xboole_0(X2,k4_xboole_0(X3,X1)))), inference(spm,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_108, plain, (k4_xboole_0(k2_xboole_0(X1,X2),k4_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[c_0_106, c_0_99])).
cnf(c_0_109, plain, (k2_xboole_0(X1,k2_xboole_0(X2,X1))=k2_xboole_0(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_46]), c_0_21])).
cnf(c_0_110, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X3,X1)))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_107]), c_0_39]), c_0_37])).
fof(c_0_111, negated_conjecture, ~(![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(assume_negation,[status(cth)],[t47_xboole_1])).
cnf(c_0_112, plain, (k4_xboole_0(k2_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_42]), c_0_109])).
cnf(c_0_113, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X2,X1)), inference(rw,[status(thm)],[c_0_91, c_0_110])).
cnf(c_0_114, plain, (k3_xboole_0(X1,k3_xboole_0(X2,k4_xboole_0(X3,k3_xboole_0(X1,X2))))=k1_xboole_0), inference(spm,[status(thm)],[c_0_39, c_0_81])).
fof(c_0_115, negated_conjecture, k4_xboole_0(esk3_0,k3_xboole_0(esk3_0,esk4_0))!=k4_xboole_0(esk3_0,esk4_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_111])])])).
cnf(c_0_116, plain, (k4_xboole_0(X1,k3_xboole_0(X2,X1))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_56])).
cnf(c_0_117, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k3_xboole_0(X2,X1)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_84, c_0_114])).
cnf(c_0_118, negated_conjecture, (k4_xboole_0(esk3_0,k3_xboole_0(esk3_0,esk4_0))!=k4_xboole_0(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_115])).
cnf(c_0_119, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_48]), c_0_20]), c_0_48]), c_0_21]), c_0_63])).
cnf(c_0_120, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_118, c_0_119])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 121
# Proof object clause steps            : 84
# Proof object formula steps           : 37
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 61
# Proof object simplifying inferences  : 28
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 6774
# ...of these trivial                  : 2329
# ...subsumed                          : 3342
# ...remaining for further processing  : 1103
# Other redundant clauses eliminated   : 1065
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 78
# Generated clauses                    : 278993
# ...of the previous two non-trivial   : 101126
# Contextual simplify-reflections      : 1
# Paramodulations                      : 277922
# Factorizations                       : 6
# Equation resolutions                 : 1065
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1016
#    Positive orientable unit clauses  : 845
#    Positive unorientable unit clauses: 5
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 163
# Current number of unprocessed clauses: 94011
# ...number of literals in the above   : 128286
# Current number of archived formulas  : 0
# Current number of archived clauses   : 84
# Clause-clause subsumption calls (NU) : 10229
# Rec. Clause-clause subsumption calls : 7696
# Non-unit clause-clause subsumptions  : 1554
# Unit Clause-clause subsumption calls : 182
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5331
# BW rewrite match successes           : 159
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2080876

# -------------------------------------------------
# User time                : 1.229 s
# System time              : 0.030 s
# Total time               : 1.259 s
# Maximum resident set size: 86916 pages
