# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t62_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1)))=>![X3, X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(X5=X4=>((r1_lattice3(X1,X3,X4)=>r1_lattice3(X2,X3,X5))&(r2_lattice3(X1,X3,X4)=>r2_lattice3(X2,X3,X5)))))))), file('data/MPTP/proofs/56986b2e68704232a296dda969ecc079.E_input', t62_yellow_0)).
fof(t61_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:((v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>((X5=X3&(X6=X4&(r1_orders_2(X1,X3,X4)&r2_hidden(X5,u1_struct_0(X2)))))=>r1_orders_2(X2,X5,X6)))))))), file('data/MPTP/proofs/56986b2e68704232a296dda969ecc079.E_input', t61_yellow_0)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/56986b2e68704232a296dda969ecc079.E_input', d9_lattice3)).
fof(dt_m1_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_yellow_0(X2,X1)=>l1_orders_2(X2))), file('data/MPTP/proofs/56986b2e68704232a296dda969ecc079.E_input', dt_m1_yellow_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/56986b2e68704232a296dda969ecc079.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/56986b2e68704232a296dda969ecc079.E_input', dt_l1_orders_2)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/56986b2e68704232a296dda969ecc079.E_input', d2_subset_1)).
fof(t59_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((~(v2_struct_0(X2))&m1_yellow_0(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>m1_subset_1(X3,u1_struct_0(X1))))), file('data/MPTP/proofs/56986b2e68704232a296dda969ecc079.E_input', t59_yellow_0)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/56986b2e68704232a296dda969ecc079.E_input', d8_lattice3)).
fof(c_0_9, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1)))=>![X3, X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(X5=X4=>((r1_lattice3(X1,X3,X4)=>r1_lattice3(X2,X3,X5))&(r2_lattice3(X1,X3,X4)=>r2_lattice3(X2,X3,X5))))))))), inference(assume_negation,[status(cth)],[t62_yellow_0])).
fof(c_0_10, plain, ![X12, X13, X14, X15, X16, X17]:(~l1_orders_2(X12)|(~v4_yellow_0(X13,X12)|~m1_yellow_0(X13,X12)|(~m1_subset_1(X14,u1_struct_0(X12))|(~m1_subset_1(X15,u1_struct_0(X12))|(~m1_subset_1(X16,u1_struct_0(X13))|(~m1_subset_1(X17,u1_struct_0(X13))|(X16!=X14|(X17!=X15|(~r1_orders_2(X12,X14,X15)|~r2_hidden(X16,u1_struct_0(X13))))|r1_orders_2(X13,X16,X17)))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t61_yellow_0])])])).
fof(c_0_11, plain, ![X19, X20, X21, X22]:((~r2_lattice3(X19,X20,X21)|(~m1_subset_1(X22,u1_struct_0(X19))|(~r2_hidden(X22,X20)|r1_orders_2(X19,X22,X21)))|~m1_subset_1(X21,u1_struct_0(X19))|~l1_orders_2(X19))&((m1_subset_1(esk6_3(X19,X20,X21),u1_struct_0(X19))|r2_lattice3(X19,X20,X21)|~m1_subset_1(X21,u1_struct_0(X19))|~l1_orders_2(X19))&((r2_hidden(esk6_3(X19,X20,X21),X20)|r2_lattice3(X19,X20,X21)|~m1_subset_1(X21,u1_struct_0(X19))|~l1_orders_2(X19))&(~r1_orders_2(X19,esk6_3(X19,X20,X21),X21)|r2_lattice3(X19,X20,X21)|~m1_subset_1(X21,u1_struct_0(X19))|~l1_orders_2(X19))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
fof(c_0_12, negated_conjecture, ((~v2_struct_0(esk1_0)&l1_orders_2(esk1_0))&((~v2_struct_0(esk2_0)&(v4_yellow_0(esk2_0,esk1_0)&m1_yellow_0(esk2_0,esk1_0)))&(m1_subset_1(esk4_0,u1_struct_0(esk1_0))&(m1_subset_1(esk5_0,u1_struct_0(esk2_0))&(esk5_0=esk4_0&(((r2_lattice3(esk1_0,esk3_0,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0))&(~r2_lattice3(esk2_0,esk3_0,esk5_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)))&((r2_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk2_0,esk3_0,esk5_0))&(~r2_lattice3(esk2_0,esk3_0,esk5_0)|~r1_lattice3(esk2_0,esk3_0,esk5_0))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])])).
cnf(c_0_13, plain, (r1_orders_2(X2,X5,X6)|~l1_orders_2(X1)|~v4_yellow_0(X2,X1)|~m1_yellow_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|X5!=X3|X6!=X4|~r1_orders_2(X1,X3,X4)|~r2_hidden(X5,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (r1_orders_2(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, plain, (r1_orders_2(X1,X2,X3)|~r2_hidden(X2,u1_struct_0(X1))|~r1_orders_2(X4,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))|~m1_yellow_0(X1,X4)|~v4_yellow_0(X1,X4)|~l1_orders_2(X4)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_13])])).
cnf(c_0_19, negated_conjecture, (r1_orders_2(esk1_0,X1,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_17])])).
cnf(c_0_20, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, negated_conjecture, (esk5_0=esk4_0), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_22, plain, ![X25, X26]:(~l1_orders_2(X25)|(~m1_yellow_0(X26,X25)|l1_orders_2(X26))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_yellow_0])])])).
cnf(c_0_23, negated_conjecture, (r1_orders_2(X1,X2,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_hidden(X2,u1_struct_0(X1))|~r2_hidden(X2,esk3_0)|~m1_subset_1(esk4_0,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_yellow_0(X1,esk1_0)|~v4_yellow_0(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_16]), c_0_17])])).
cnf(c_0_24, negated_conjecture, (v4_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_25, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_26, negated_conjecture, (m1_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_27, plain, (l1_orders_2(X2)|~l1_orders_2(X1)|~m1_yellow_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_28, plain, ![X18]:(v2_struct_0(X18)|~l1_struct_0(X18)|~v1_xboole_0(u1_struct_0(X18))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_29, plain, ![X24]:(~l1_orders_2(X24)|l1_struct_0(X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_30, plain, (r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk6_3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_31, negated_conjecture, (r1_orders_2(esk2_0,X1,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25]), c_0_26])])).
cnf(c_0_32, negated_conjecture, (l1_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_26]), c_0_17])])).
cnf(c_0_33, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_lattice3(esk2_0,esk3_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_34, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk2_0,esk3_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_37, negated_conjecture, (r2_lattice3(esk2_0,X1,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_hidden(esk6_3(esk2_0,X1,esk4_0),u1_struct_0(esk2_0))|~r2_hidden(esk6_3(esk2_0,X1,esk4_0),esk3_0)|~m1_subset_1(esk6_3(esk2_0,X1,esk4_0),u1_struct_0(esk1_0))|~m1_subset_1(esk6_3(esk2_0,X1,esk4_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_25]), c_0_32])])).
cnf(c_0_38, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_39, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_lattice3(esk2_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_33, c_0_21])).
fof(c_0_40, plain, ![X27, X28]:(((~m1_subset_1(X28,X27)|r2_hidden(X28,X27)|v1_xboole_0(X27))&(~r2_hidden(X28,X27)|m1_subset_1(X28,X27)|v1_xboole_0(X27)))&((~m1_subset_1(X28,X27)|v1_xboole_0(X28)|~v1_xboole_0(X27))&(~v1_xboole_0(X28)|m1_subset_1(X28,X27)|~v1_xboole_0(X27)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
cnf(c_0_41, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_42, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
fof(c_0_43, plain, ![X29, X30, X31]:(v2_struct_0(X29)|~l1_orders_2(X29)|(v2_struct_0(X30)|~m1_yellow_0(X30,X29)|(~m1_subset_1(X31,u1_struct_0(X30))|m1_subset_1(X31,u1_struct_0(X29))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t59_yellow_0])])])])).
cnf(c_0_44, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk2_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_36, c_0_21])).
cnf(c_0_45, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_hidden(esk6_3(esk2_0,esk3_0,esk4_0),u1_struct_0(esk2_0))|~m1_subset_1(esk6_3(esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0))|~m1_subset_1(esk6_3(esk2_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_25]), c_0_32])]), c_0_39])).
cnf(c_0_46, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_47, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_32])])).
cnf(c_0_48, plain, (v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~m1_yellow_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_50, negated_conjecture, (r1_orders_2(esk1_0,X1,esk4_0)|~r2_hidden(X1,esk3_0)|~r1_lattice3(esk2_0,esk3_0,esk4_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_44]), c_0_16]), c_0_17])])).
cnf(c_0_51, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk4_0)|~m1_subset_1(esk6_3(esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0))|~m1_subset_1(esk6_3(esk2_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])).
cnf(c_0_52, plain, (r2_lattice3(X1,X2,X3)|m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X4))|v2_struct_0(X4)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_yellow_0(X1,X4)|~l1_orders_2(X4)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_27])).
cnf(c_0_53, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_54, negated_conjecture, (r1_orders_2(X1,X2,esk4_0)|~r2_hidden(X2,u1_struct_0(X1))|~r2_hidden(X2,esk3_0)|~r1_lattice3(esk2_0,esk3_0,esk4_0)|~m1_subset_1(esk4_0,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_yellow_0(X1,esk1_0)|~v4_yellow_0(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_50]), c_0_16]), c_0_17])])).
fof(c_0_55, plain, ![X32, X33, X34, X35]:((~r1_lattice3(X32,X33,X34)|(~m1_subset_1(X35,u1_struct_0(X32))|(~r2_hidden(X35,X33)|r1_orders_2(X32,X34,X35)))|~m1_subset_1(X34,u1_struct_0(X32))|~l1_orders_2(X32))&((m1_subset_1(esk7_3(X32,X33,X34),u1_struct_0(X32))|r1_lattice3(X32,X33,X34)|~m1_subset_1(X34,u1_struct_0(X32))|~l1_orders_2(X32))&((r2_hidden(esk7_3(X32,X33,X34),X33)|r1_lattice3(X32,X33,X34)|~m1_subset_1(X34,u1_struct_0(X32))|~l1_orders_2(X32))&(~r1_orders_2(X32,X34,esk7_3(X32,X33,X34))|r1_lattice3(X32,X33,X34)|~m1_subset_1(X34,u1_struct_0(X32))|~l1_orders_2(X32))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
cnf(c_0_56, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk4_0)|~m1_subset_1(esk6_3(esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_49]), c_0_25]), c_0_32])]), c_0_39])).
cnf(c_0_57, negated_conjecture, (r2_lattice3(esk2_0,X1,X2)|m1_subset_1(esk6_3(esk2_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_26]), c_0_17])]), c_0_53]), c_0_41])).
cnf(c_0_58, negated_conjecture, (r1_orders_2(esk2_0,X1,esk4_0)|~r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk3_0)|~r1_lattice3(esk2_0,esk3_0,esk4_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_24]), c_0_25]), c_0_26])])).
cnf(c_0_59, negated_conjecture, (~r2_lattice3(esk2_0,esk3_0,esk5_0)|~r1_lattice3(esk2_0,esk3_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_60, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_61, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_25])]), c_0_39])).
cnf(c_0_62, negated_conjecture, (r2_lattice3(esk2_0,X1,esk4_0)|~r2_hidden(esk6_3(esk2_0,X1,esk4_0),u1_struct_0(esk2_0))|~r2_hidden(esk6_3(esk2_0,X1,esk4_0),esk3_0)|~r1_lattice3(esk2_0,esk3_0,esk4_0)|~m1_subset_1(esk6_3(esk2_0,X1,esk4_0),u1_struct_0(esk1_0))|~m1_subset_1(esk6_3(esk2_0,X1,esk4_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_58]), c_0_25]), c_0_32])])).
cnf(c_0_63, negated_conjecture, (~r2_lattice3(esk2_0,esk3_0,esk4_0)|~r1_lattice3(esk2_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_21]), c_0_21])).
cnf(c_0_64, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,X1)|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_16]), c_0_17])])).
cnf(c_0_65, negated_conjecture, (~r2_hidden(esk6_3(esk2_0,esk3_0,esk4_0),u1_struct_0(esk2_0))|~r1_lattice3(esk2_0,esk3_0,esk4_0)|~m1_subset_1(esk6_3(esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0))|~m1_subset_1(esk6_3(esk2_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_38]), c_0_25]), c_0_32])]), c_0_63])).
cnf(c_0_66, negated_conjecture, (r1_orders_2(X1,esk4_0,X2)|~r2_hidden(esk4_0,u1_struct_0(X1))|~r2_hidden(X2,esk3_0)|~m1_subset_1(esk4_0,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_yellow_0(X1,esk1_0)|~v4_yellow_0(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_64]), c_0_16]), c_0_17])])).
cnf(c_0_67, negated_conjecture, (~r1_lattice3(esk2_0,esk3_0,esk4_0)|~m1_subset_1(esk6_3(esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0))|~m1_subset_1(esk6_3(esk2_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_46]), c_0_47])).
cnf(c_0_68, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk7_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_69, negated_conjecture, (r1_orders_2(esk2_0,esk4_0,X1)|~r2_hidden(esk4_0,u1_struct_0(esk2_0))|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_24]), c_0_25]), c_0_26])])).
cnf(c_0_70, negated_conjecture, (~r1_lattice3(esk2_0,esk3_0,esk4_0)|~m1_subset_1(esk6_3(esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_49]), c_0_25]), c_0_32])]), c_0_63])).
cnf(c_0_71, negated_conjecture, (r1_lattice3(esk2_0,X1,esk4_0)|~r2_hidden(esk7_3(esk2_0,X1,esk4_0),esk3_0)|~r2_hidden(esk4_0,u1_struct_0(esk2_0))|~m1_subset_1(esk7_3(esk2_0,X1,esk4_0),u1_struct_0(esk1_0))|~m1_subset_1(esk7_3(esk2_0,X1,esk4_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_25]), c_0_32])])).
cnf(c_0_72, plain, (r2_hidden(esk7_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_73, negated_conjecture, (~r1_lattice3(esk2_0,esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_57]), c_0_25])]), c_0_63])).
cnf(c_0_74, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_75, negated_conjecture, (~r2_hidden(esk4_0,u1_struct_0(esk2_0))|~m1_subset_1(esk7_3(esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0))|~m1_subset_1(esk7_3(esk2_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_25]), c_0_32])]), c_0_73])).
cnf(c_0_76, plain, (r1_lattice3(X1,X2,X3)|m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X4))|v2_struct_0(X4)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_yellow_0(X1,X4)|~l1_orders_2(X4)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_74]), c_0_27])).
cnf(c_0_77, negated_conjecture, (~r2_hidden(esk4_0,u1_struct_0(esk2_0))|~m1_subset_1(esk7_3(esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_74]), c_0_25]), c_0_32])]), c_0_73])).
cnf(c_0_78, negated_conjecture, (r1_lattice3(esk2_0,X1,X2)|m1_subset_1(esk7_3(esk2_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_26]), c_0_17])]), c_0_53]), c_0_41])).
cnf(c_0_79, negated_conjecture, (~r2_hidden(esk4_0,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_25])]), c_0_73])).
cnf(c_0_80, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_46]), c_0_25])]), c_0_47]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 81
# Proof object clause steps            : 62
# Proof object formula steps           : 19
# Proof object conjectures             : 47
# Proof object clause conjectures      : 44
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 31
# Proof object simplifying inferences  : 95
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 29
# Processed clauses                    : 126
# ...of these trivial                  : 0
# ...subsumed                          : 14
# ...remaining for further processing  : 112
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 9
# Generated clauses                    : 85
# ...of the previous two non-trivial   : 78
# Contextual simplify-reflections      : 8
# Paramodulations                      : 84
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 68
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 54
# Current number of unprocessed clauses: 10
# ...number of literals in the above   : 87
# Current number of archived formulas  : 0
# Current number of archived clauses   : 43
# Clause-clause subsumption calls (NU) : 2035
# Rec. Clause-clause subsumption calls : 273
# Non-unit clause-clause subsumptions  : 17
# Unit Clause-clause subsumption calls : 34
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5148

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.000 s
# Total time               : 0.013 s
# Maximum resident set size: 28992 pages
