# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>(r2_hidden(X1,k1_relat_1(X3))&X2=k1_funct_1(X3,X1)))), file('data/MPTP/proofs/56aa46223a944b719275ecb2c3b149cb.E_input', t8_funct_1)).
fof(t116_funct_2, conjecture, ![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X5]:~((r2_hidden(X5,k7_relset_1(X1,X2,X4,X3))&![X6]:(m1_subset_1(X6,X1)=>~((r2_hidden(X6,X3)&X5=k1_funct_1(X4,X6))))))), file('data/MPTP/proofs/56aa46223a944b719275ecb2c3b149cb.E_input', t116_funct_2)).
fof(s3_funct_1__e9_44_1__funct_1, axiom, ![X1]:?[X2]:((v1_relat_1(X2)&v1_funct_1(X2))&(k1_relat_1(X2)=X1&![X3]:(r2_hidden(X3,X1)=>k1_funct_1(X2,X3)=k1_xboole_0))), file('data/MPTP/proofs/56aa46223a944b719275ecb2c3b149cb.E_input', s3_funct_1__e9_44_1__funct_1)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/56aa46223a944b719275ecb2c3b149cb.E_input', l3_subset_1)).
fof(t143_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k9_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k1_relat_1(X3))&(r2_hidden(k4_tarski(X4,X1),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/56aa46223a944b719275ecb2c3b149cb.E_input', t143_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/56aa46223a944b719275ecb2c3b149cb.E_input', fc6_relat_1)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/56aa46223a944b719275ecb2c3b149cb.E_input', cc2_relat_1)).
fof(t64_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>((k1_relat_1(X1)=k1_xboole_0|k2_relat_1(X1)=k1_xboole_0)=>X1=k1_xboole_0)), file('data/MPTP/proofs/56aa46223a944b719275ecb2c3b149cb.E_input', t64_relat_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/56aa46223a944b719275ecb2c3b149cb.E_input', d1_xboole_0)).
fof(redefinition_k7_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k7_relset_1(X1,X2,X3,X4)=k9_relat_1(X3,X4)), file('data/MPTP/proofs/56aa46223a944b719275ecb2c3b149cb.E_input', redefinition_k7_relset_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/56aa46223a944b719275ecb2c3b149cb.E_input', redefinition_k1_relset_1)).
fof(cc1_relat_1, axiom, ![X1]:(v1_xboole_0(X1)=>v1_relat_1(X1)), file('data/MPTP/proofs/56aa46223a944b719275ecb2c3b149cb.E_input', cc1_relat_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/56aa46223a944b719275ecb2c3b149cb.E_input', t1_subset)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(((X2=k1_xboole_0=>X1=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X2,X3)))&(X2=k1_xboole_0=>(X1=k1_xboole_0|(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0))))), file('data/MPTP/proofs/56aa46223a944b719275ecb2c3b149cb.E_input', d1_funct_2)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/56aa46223a944b719275ecb2c3b149cb.E_input', fc1_xboole_0)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/56aa46223a944b719275ecb2c3b149cb.E_input', t113_zfmisc_1)).
fof(c_0_16, plain, ![X7, X8, X9]:(((r2_hidden(X7,k1_relat_1(X9))|~r2_hidden(k4_tarski(X7,X8),X9)|(~v1_relat_1(X9)|~v1_funct_1(X9)))&(X8=k1_funct_1(X9,X7)|~r2_hidden(k4_tarski(X7,X8),X9)|(~v1_relat_1(X9)|~v1_funct_1(X9))))&(~r2_hidden(X7,k1_relat_1(X9))|X8!=k1_funct_1(X9,X7)|r2_hidden(k4_tarski(X7,X8),X9)|(~v1_relat_1(X9)|~v1_funct_1(X9)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_funct_1])])])).
fof(c_0_17, negated_conjecture, ~(![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X5]:~((r2_hidden(X5,k7_relset_1(X1,X2,X4,X3))&![X6]:(m1_subset_1(X6,X1)=>~((r2_hidden(X6,X3)&X5=k1_funct_1(X4,X6)))))))), inference(assume_negation,[status(cth)],[t116_funct_2])).
cnf(c_0_18, plain, (r2_hidden(k4_tarski(X1,X3),X2)|~r2_hidden(X1,k1_relat_1(X2))|X3!=k1_funct_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_19, plain, ![X16, X18]:((v1_relat_1(esk6_1(X16))&v1_funct_1(esk6_1(X16)))&(k1_relat_1(esk6_1(X16))=X16&(~r2_hidden(X18,X16)|k1_funct_1(esk6_1(X16),X18)=k1_xboole_0))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[s3_funct_1__e9_44_1__funct_1])])])])).
fof(c_0_20, plain, ![X19, X20, X21]:(~m1_subset_1(X20,k1_zfmisc_1(X19))|(~r2_hidden(X21,X20)|r2_hidden(X21,X19))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
fof(c_0_21, negated_conjecture, ![X15]:((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,esk1_0,esk2_0)&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))&(r2_hidden(esk5_0,k7_relset_1(esk1_0,esk2_0,esk4_0,esk3_0))&(~m1_subset_1(X15,esk1_0)|(~r2_hidden(X15,esk3_0)|esk5_0!=k1_funct_1(esk4_0,X15))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])])).
fof(c_0_22, plain, ![X22, X23, X24, X26]:(((r2_hidden(esk7_3(X22,X23,X24),k1_relat_1(X24))|~r2_hidden(X22,k9_relat_1(X24,X23))|~v1_relat_1(X24))&((r2_hidden(k4_tarski(esk7_3(X22,X23,X24),X22),X24)|~r2_hidden(X22,k9_relat_1(X24,X23))|~v1_relat_1(X24))&(r2_hidden(esk7_3(X22,X23,X24),X23)|~r2_hidden(X22,k9_relat_1(X24,X23))|~v1_relat_1(X24))))&(~r2_hidden(X26,k1_relat_1(X24))|(~r2_hidden(k4_tarski(X26,X22),X24)|~r2_hidden(X26,X23))|r2_hidden(X22,k9_relat_1(X24,X23))|~v1_relat_1(X24))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t143_relat_1])])])])])).
cnf(c_0_23, plain, (r2_hidden(k4_tarski(X1,k1_funct_1(X2,X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (k1_funct_1(esk6_1(X2),X1)=k1_xboole_0|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (k1_relat_1(esk6_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (v1_funct_1(esk6_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (v1_relat_1(esk6_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_30, plain, ![X27, X28]:v1_relat_1(k2_zfmisc_1(X27,X28)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_31, plain, ![X29, X30]:(~v1_relat_1(X29)|(~m1_subset_1(X30,k1_zfmisc_1(X29))|v1_relat_1(X30))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
cnf(c_0_32, plain, (r2_hidden(X3,k9_relat_1(X2,X4))|~r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)|~r2_hidden(X1,X4)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33, plain, (r2_hidden(k4_tarski(X1,k1_xboole_0),esk6_1(X2))|~r2_hidden(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25]), c_0_26]), c_0_27])])).
cnf(c_0_34, plain, (X1=k1_funct_1(X2,X3)|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_35, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(esk1_0,esk2_0))|~r2_hidden(X1,esk4_0)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_36, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_39, plain, ![X31]:((k1_relat_1(X31)!=k1_xboole_0|X31=k1_xboole_0|~v1_relat_1(X31))&(k2_relat_1(X31)!=k1_xboole_0|X31=k1_xboole_0|~v1_relat_1(X31))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t64_relat_1])])])).
cnf(c_0_40, plain, (r2_hidden(k1_xboole_0,k9_relat_1(esk6_1(X1),X2))|~r2_hidden(X3,X1)|~r2_hidden(X3,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_25]), c_0_27])])).
cnf(c_0_41, negated_conjecture, (X1=k1_funct_1(k2_zfmisc_1(esk1_0,esk2_0),X2)|~r2_hidden(k4_tarski(X2,X1),esk4_0)|~v1_funct_1(k2_zfmisc_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])])).
cnf(c_0_42, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_43, negated_conjecture, (v1_relat_1(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_29]), c_0_36])])).
cnf(c_0_44, negated_conjecture, (r2_hidden(X1,k1_relat_1(k2_zfmisc_1(esk1_0,esk2_0)))|~r2_hidden(k4_tarski(X1,X2),esk4_0)|~v1_funct_1(k2_zfmisc_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_35]), c_0_36])])).
fof(c_0_45, plain, ![X32, X33, X34]:((~v1_xboole_0(X32)|~r2_hidden(X33,X32))&(r2_hidden(esk8_1(X34),X34)|v1_xboole_0(X34))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_46, plain, ![X36, X37, X38, X39]:(~m1_subset_1(X38,k1_zfmisc_1(k2_zfmisc_1(X36,X37)))|k7_relset_1(X36,X37,X38,X39)=k9_relat_1(X38,X39)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_relset_1])])).
cnf(c_0_47, plain, (r2_hidden(k4_tarski(esk7_3(X1,X2,X3),X1),X3)|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_48, plain, ![X40, X41, X42]:(~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(X40,X41)))|k1_relset_1(X40,X41,X42)=k1_relat_1(X42)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
fof(c_0_49, plain, ![X43]:(~v1_xboole_0(X43)|v1_relat_1(X43)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relat_1])])).
cnf(c_0_50, plain, (X1=k1_xboole_0|k1_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51, negated_conjecture, (r2_hidden(k1_xboole_0,k9_relat_1(esk6_1(k2_zfmisc_1(esk1_0,esk2_0)),X1))|~r2_hidden(X2,esk4_0)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_40, c_0_35])).
cnf(c_0_52, negated_conjecture, (k1_funct_1(k2_zfmisc_1(esk1_0,esk2_0),X1)=k1_funct_1(esk4_0,X1)|~r2_hidden(X1,k1_relat_1(esk4_0))|~v1_funct_1(k2_zfmisc_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_23]), c_0_42]), c_0_43])])).
cnf(c_0_53, negated_conjecture, (r2_hidden(X1,k1_relat_1(k2_zfmisc_1(esk1_0,esk2_0)))|~r2_hidden(X1,k1_relat_1(esk4_0))|~v1_funct_1(k2_zfmisc_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_23]), c_0_42]), c_0_43])])).
cnf(c_0_54, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_55, plain, (r2_hidden(esk7_3(X1,X2,X3),k1_relat_1(X3))|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_56, negated_conjecture, (r2_hidden(esk5_0,k7_relset_1(esk1_0,esk2_0,esk4_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_57, plain, (k7_relset_1(X2,X3,X1,X4)=k9_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_58, negated_conjecture, (~m1_subset_1(X1,esk1_0)|~r2_hidden(X1,esk3_0)|esk5_0!=k1_funct_1(esk4_0,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_59, plain, (k1_funct_1(X1,esk7_3(X2,X3,X1))=X2|~r2_hidden(X2,k9_relat_1(X1,X3))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_34, c_0_47])).
fof(c_0_60, plain, ![X44, X45]:(~r2_hidden(X44,X45)|m1_subset_1(X44,X45)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
fof(c_0_61, plain, ![X46, X47, X48]:((((~v1_funct_2(X48,X46,X47)|X46=k1_relset_1(X46,X47,X48)|X47=k1_xboole_0|~m1_subset_1(X48,k1_zfmisc_1(k2_zfmisc_1(X46,X47))))&(X46!=k1_relset_1(X46,X47,X48)|v1_funct_2(X48,X46,X47)|X47=k1_xboole_0|~m1_subset_1(X48,k1_zfmisc_1(k2_zfmisc_1(X46,X47)))))&((~v1_funct_2(X48,X46,X47)|X46=k1_relset_1(X46,X47,X48)|X46!=k1_xboole_0|~m1_subset_1(X48,k1_zfmisc_1(k2_zfmisc_1(X46,X47))))&(X46!=k1_relset_1(X46,X47,X48)|v1_funct_2(X48,X46,X47)|X46!=k1_xboole_0|~m1_subset_1(X48,k1_zfmisc_1(k2_zfmisc_1(X46,X47))))))&((~v1_funct_2(X48,X46,X47)|X48=k1_xboole_0|X46=k1_xboole_0|X47!=k1_xboole_0|~m1_subset_1(X48,k1_zfmisc_1(k2_zfmisc_1(X46,X47))))&(X48!=k1_xboole_0|v1_funct_2(X48,X46,X47)|X46=k1_xboole_0|X47!=k1_xboole_0|~m1_subset_1(X48,k1_zfmisc_1(k2_zfmisc_1(X46,X47)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
cnf(c_0_62, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_63, plain, (v1_relat_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_64, plain, (esk6_1(k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_25]), c_0_27])])])).
cnf(c_0_65, negated_conjecture, (r2_hidden(k1_xboole_0,k9_relat_1(esk6_1(k2_zfmisc_1(esk1_0,esk2_0)),X1))|~r2_hidden(k4_tarski(X2,k1_funct_1(esk4_0,X2)),X1)|~r2_hidden(X2,k1_relat_1(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_23]), c_0_42]), c_0_43])])).
cnf(c_0_66, negated_conjecture, (r2_hidden(k4_tarski(X1,k1_funct_1(esk4_0,X1)),k2_zfmisc_1(esk1_0,esk2_0))|~r2_hidden(X1,k1_relat_1(esk4_0))|~v1_funct_1(k2_zfmisc_1(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_52]), c_0_36])]), c_0_53])).
cnf(c_0_67, plain, (~v1_xboole_0(k1_relat_1(X1))|~r2_hidden(X2,k9_relat_1(X1,X3))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_68, negated_conjecture, (r2_hidden(esk5_0,k9_relat_1(esk4_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_29])])).
cnf(c_0_69, negated_conjecture, (~m1_subset_1(esk7_3(esk5_0,X1,esk4_0),esk1_0)|~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk3_0)|~r2_hidden(esk5_0,k9_relat_1(esk4_0,X1))), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_42]), c_0_43])])])).
cnf(c_0_70, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_71, plain, (X2=k1_relset_1(X2,X3,X1)|X3=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_72, negated_conjecture, (k1_relset_1(esk1_0,esk2_0,esk4_0)=k1_relat_1(esk4_0)), inference(spm,[status(thm)],[c_0_62, c_0_29])).
cnf(c_0_73, negated_conjecture, (v1_funct_2(esk4_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_74, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k1_relat_1(X1))|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_23]), c_0_63])).
cnf(c_0_75, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_25, c_0_64])).
cnf(c_0_76, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_77, plain, (v1_funct_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_26, c_0_64])).
cnf(c_0_78, negated_conjecture, (r2_hidden(k1_xboole_0,k9_relat_1(esk6_1(k2_zfmisc_1(esk1_0,esk2_0)),k2_zfmisc_1(esk1_0,esk2_0)))|~r2_hidden(X1,k1_relat_1(esk4_0))|~v1_funct_1(k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_79, plain, (r2_hidden(esk8_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_80, negated_conjecture, (~v1_xboole_0(k1_relat_1(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_43])])).
cnf(c_0_81, negated_conjecture, (~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk3_0)|~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk1_0)|~r2_hidden(esk5_0,k9_relat_1(esk4_0,X1))), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_82, plain, (r2_hidden(esk7_3(X1,X2,X3),X2)|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_83, negated_conjecture, (k1_relat_1(esk4_0)=esk1_0|esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_29]), c_0_72]), c_0_73])])).
fof(c_0_84, plain, ![X49, X50]:((k2_zfmisc_1(X49,X50)!=k1_xboole_0|(X49=k1_xboole_0|X50=k1_xboole_0))&((X49!=k1_xboole_0|k2_zfmisc_1(X49,X50)=k1_xboole_0)&(X50!=k1_xboole_0|k2_zfmisc_1(X49,X50)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_85, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76]), c_0_77])])).
cnf(c_0_86, plain, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_63, c_0_76])).
cnf(c_0_87, negated_conjecture, (r2_hidden(k1_xboole_0,k9_relat_1(esk6_1(k2_zfmisc_1(esk1_0,esk2_0)),k2_zfmisc_1(esk1_0,esk2_0)))|~v1_funct_1(k2_zfmisc_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80])).
cnf(c_0_88, negated_conjecture, (~r2_hidden(esk7_3(esk5_0,esk3_0,esk4_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_68]), c_0_43])])).
cnf(c_0_89, negated_conjecture, (esk2_0=k1_xboole_0|r2_hidden(esk7_3(X1,X2,esk4_0),esk1_0)|~r2_hidden(X1,k9_relat_1(esk4_0,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_83]), c_0_43])])).
cnf(c_0_90, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_91, plain, (~r2_hidden(X1,k9_relat_1(k1_xboole_0,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_47]), c_0_86])])).
cnf(c_0_92, negated_conjecture, (~v1_xboole_0(k9_relat_1(esk6_1(k2_zfmisc_1(esk1_0,esk2_0)),k2_zfmisc_1(esk1_0,esk2_0)))|~v1_funct_1(k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_54, c_0_87])).
cnf(c_0_93, negated_conjecture, (esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_68])])).
cnf(c_0_94, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_90])).
cnf(c_0_95, plain, (v1_xboole_0(k9_relat_1(k1_xboole_0,X1))), inference(spm,[status(thm)],[c_0_91, c_0_79])).
cnf(c_0_96, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_93]), c_0_94]), c_0_64]), c_0_93]), c_0_94]), c_0_95]), c_0_93]), c_0_94]), c_0_77])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 97
# Proof object clause steps            : 65
# Proof object formula steps           : 32
# Proof object conjectures             : 30
# Proof object clause conjectures      : 27
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 29
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 33
# Proof object simplifying inferences  : 65
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 37
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 37
# Processed clauses                    : 665
# ...of these trivial                  : 0
# ...subsumed                          : 288
# ...remaining for further processing  : 377
# Other redundant clauses eliminated   : 10
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 53
# Generated clauses                    : 1639
# ...of the previous two non-trivial   : 1605
# Contextual simplify-reflections      : 22
# Paramodulations                      : 1630
# Factorizations                       : 0
# Equation resolutions                 : 10
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 276
#    Positive orientable unit clauses  : 25
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 15
#    Non-unit-clauses                  : 236
# Current number of unprocessed clauses: 995
# ...number of literals in the above   : 4678
# Current number of archived formulas  : 0
# Current number of archived clauses   : 94
# Clause-clause subsumption calls (NU) : 24208
# Rec. Clause-clause subsumption calls : 8891
# Non-unit clause-clause subsumptions  : 197
# Unit Clause-clause subsumption calls : 676
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 27
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 34827

# -------------------------------------------------
# User time                : 0.032 s
# System time              : 0.003 s
# Total time               : 0.035 s
# Maximum resident set size: 28740 pages
