# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t90_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((~(v2_struct_0(X5))&m1_pre_topc(X5,X1))=>((m1_pre_topc(X3,X4)&m1_pre_topc(X5,X3))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))=>((v1_funct_1(k3_tmap_1(X1,X2,X4,X3,X6))&(v1_funct_2(k3_tmap_1(X1,X2,X4,X3,X6),u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(k3_tmap_1(X1,X2,X4,X3,X6),X3,X2)&m1_subset_1(k3_tmap_1(X1,X2,X4,X3,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))=>(v1_funct_1(k3_tmap_1(X1,X2,X4,X5,X6))&(v1_funct_2(k3_tmap_1(X1,X2,X4,X5,X6),u1_struct_0(X5),u1_struct_0(X2))&(v5_pre_topc(k3_tmap_1(X1,X2,X4,X5,X6),X5,X2)&m1_subset_1(k3_tmap_1(X1,X2,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X2))))))))))))))), file('data/MPTP/proofs/56e4d87800754e529f18bd50da57131b.E_input', t90_tmap_1)).
fof(t79_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((~(v2_struct_0(X5))&m1_pre_topc(X5,X1))=>((m1_pre_topc(X4,X3)&m1_pre_topc(X5,X4))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))=>r2_funct_2(u1_struct_0(X5),u1_struct_0(X2),k3_tmap_1(X1,X2,X4,X5,k3_tmap_1(X1,X2,X3,X4,X6)),k3_tmap_1(X1,X2,X3,X5,X6))))))))), file('data/MPTP/proofs/56e4d87800754e529f18bd50da57131b.E_input', t79_tmap_1)).
fof(dt_k3_tmap_1, axiom, ![X1, X2, X3, X4, X5]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))&(m1_pre_topc(X3,X1)&(m1_pre_topc(X4,X1)&(v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))))))=>(v1_funct_1(k3_tmap_1(X1,X2,X3,X4,X5))&(v1_funct_2(k3_tmap_1(X1,X2,X3,X4,X5),u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(k3_tmap_1(X1,X2,X3,X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))), file('data/MPTP/proofs/56e4d87800754e529f18bd50da57131b.E_input', dt_k3_tmap_1)).
fof(t89_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))=>(m1_pre_topc(X4,X3)=>(v1_funct_1(k3_tmap_1(X1,X2,X3,X4,X5))&(v1_funct_2(k3_tmap_1(X1,X2,X3,X4,X5),u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(k3_tmap_1(X1,X2,X3,X4,X5),X4,X2)&m1_subset_1(k3_tmap_1(X1,X2,X3,X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))))))))), file('data/MPTP/proofs/56e4d87800754e529f18bd50da57131b.E_input', t89_tmap_1)).
fof(redefinition_r2_funct_2, axiom, ![X1, X2, X3, X4]:(((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&(v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))=>(r2_funct_2(X1,X2,X3,X4)<=>X3=X4)), file('data/MPTP/proofs/56e4d87800754e529f18bd50da57131b.E_input', redefinition_r2_funct_2)).
fof(c_0_5, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((~(v2_struct_0(X5))&m1_pre_topc(X5,X1))=>((m1_pre_topc(X3,X4)&m1_pre_topc(X5,X3))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))=>((v1_funct_1(k3_tmap_1(X1,X2,X4,X3,X6))&(v1_funct_2(k3_tmap_1(X1,X2,X4,X3,X6),u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(k3_tmap_1(X1,X2,X4,X3,X6),X3,X2)&m1_subset_1(k3_tmap_1(X1,X2,X4,X3,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))=>(v1_funct_1(k3_tmap_1(X1,X2,X4,X5,X6))&(v1_funct_2(k3_tmap_1(X1,X2,X4,X5,X6),u1_struct_0(X5),u1_struct_0(X2))&(v5_pre_topc(k3_tmap_1(X1,X2,X4,X5,X6),X5,X2)&m1_subset_1(k3_tmap_1(X1,X2,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X2)))))))))))))))), inference(assume_negation,[status(cth)],[t90_tmap_1])).
fof(c_0_6, plain, ![X17, X18, X19, X20, X21, X22]:(v2_struct_0(X17)|(~v2_pre_topc(X17)|~l1_pre_topc(X17))|(v2_struct_0(X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18))|(v2_struct_0(X19)|~m1_pre_topc(X19,X17)|(v2_struct_0(X20)|~m1_pre_topc(X20,X17)|(v2_struct_0(X21)|~m1_pre_topc(X21,X17)|(~m1_pre_topc(X20,X19)|~m1_pre_topc(X21,X20)|(~v1_funct_1(X22)|(~v1_funct_2(X22,u1_struct_0(X19),u1_struct_0(X18))|~m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X19),u1_struct_0(X18)))))|r2_funct_2(u1_struct_0(X21),u1_struct_0(X18),k3_tmap_1(X17,X18,X20,X21,k3_tmap_1(X17,X18,X19,X20,X22)),k3_tmap_1(X17,X18,X19,X21,X22))))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t79_tmap_1])])])])).
fof(c_0_7, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk1_0))&((~v2_struct_0(esk4_0)&m1_pre_topc(esk4_0,esk1_0))&((~v2_struct_0(esk5_0)&m1_pre_topc(esk5_0,esk1_0))&((m1_pre_topc(esk3_0,esk4_0)&m1_pre_topc(esk5_0,esk3_0))&((v1_funct_1(esk6_0)&(v1_funct_2(esk6_0,u1_struct_0(esk4_0),u1_struct_0(esk2_0))&m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))))&((v1_funct_1(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0))&(v1_funct_2(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0),u1_struct_0(esk3_0),u1_struct_0(esk2_0))&(v5_pre_topc(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0),esk3_0,esk2_0)&m1_subset_1(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0)))))))&(~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk5_0,esk6_0))|(~v1_funct_2(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk5_0,esk6_0),u1_struct_0(esk5_0),u1_struct_0(esk2_0))|(~v5_pre_topc(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk5_0,esk6_0),esk5_0,esk2_0)|~m1_subset_1(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(esk2_0))))))))))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_5])])])])).
fof(c_0_8, plain, ![X23, X24, X25, X26, X27]:((v1_funct_1(k3_tmap_1(X23,X24,X25,X26,X27))|(v2_struct_0(X23)|(~v2_pre_topc(X23)|~l1_pre_topc(X23))|(v2_struct_0(X24)|(~v2_pre_topc(X24)|~l1_pre_topc(X24))|(~m1_pre_topc(X25,X23)|(~m1_pre_topc(X26,X23)|(~v1_funct_1(X27)|(~v1_funct_2(X27,u1_struct_0(X25),u1_struct_0(X24))|~m1_subset_1(X27,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X25),u1_struct_0(X24)))))))))))&((v1_funct_2(k3_tmap_1(X23,X24,X25,X26,X27),u1_struct_0(X26),u1_struct_0(X24))|(v2_struct_0(X23)|(~v2_pre_topc(X23)|~l1_pre_topc(X23))|(v2_struct_0(X24)|(~v2_pre_topc(X24)|~l1_pre_topc(X24))|(~m1_pre_topc(X25,X23)|(~m1_pre_topc(X26,X23)|(~v1_funct_1(X27)|(~v1_funct_2(X27,u1_struct_0(X25),u1_struct_0(X24))|~m1_subset_1(X27,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X25),u1_struct_0(X24)))))))))))&(m1_subset_1(k3_tmap_1(X23,X24,X25,X26,X27),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X26),u1_struct_0(X24))))|(v2_struct_0(X23)|(~v2_pre_topc(X23)|~l1_pre_topc(X23))|(v2_struct_0(X24)|(~v2_pre_topc(X24)|~l1_pre_topc(X24))|(~m1_pre_topc(X25,X23)|(~m1_pre_topc(X26,X23)|(~v1_funct_1(X27)|(~v1_funct_2(X27,u1_struct_0(X25),u1_struct_0(X24))|~m1_subset_1(X27,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X25),u1_struct_0(X24))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_tmap_1])])])])).
cnf(c_0_9, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X5)|r2_funct_2(u1_struct_0(X5),u1_struct_0(X2),k3_tmap_1(X1,X2,X4,X5,k3_tmap_1(X1,X2,X3,X4,X6)),k3_tmap_1(X1,X2,X3,X5,X6))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X3)|~m1_pre_topc(X5,X4)|~v1_funct_1(X6)|~v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, negated_conjecture, (v1_funct_2(esk6_0,u1_struct_0(esk4_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, negated_conjecture, (m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, negated_conjecture, (v1_funct_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14, negated_conjecture, (v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_15, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_16, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_17, plain, (v1_funct_1(k3_tmap_1(X1,X2,X3,X4,X5))|v2_struct_0(X1)|v2_struct_0(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_18, plain, (v1_funct_2(k3_tmap_1(X1,X2,X3,X4,X5),u1_struct_0(X4),u1_struct_0(X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_19, plain, ![X28, X29, X30, X31, X32]:((v1_funct_1(k3_tmap_1(X28,X29,X30,X31,X32))|~m1_pre_topc(X31,X30)|(~v1_funct_1(X32)|(~v1_funct_2(X32,u1_struct_0(X30),u1_struct_0(X29))|(~v5_pre_topc(X32,X30,X29)|~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X30),u1_struct_0(X29)))))))|(v2_struct_0(X31)|~m1_pre_topc(X31,X28))|(v2_struct_0(X30)|~m1_pre_topc(X30,X28))|(v2_struct_0(X29)|(~v2_pre_topc(X29)|~l1_pre_topc(X29)))|(v2_struct_0(X28)|(~v2_pre_topc(X28)|~l1_pre_topc(X28))))&((v1_funct_2(k3_tmap_1(X28,X29,X30,X31,X32),u1_struct_0(X31),u1_struct_0(X29))|~m1_pre_topc(X31,X30)|(~v1_funct_1(X32)|(~v1_funct_2(X32,u1_struct_0(X30),u1_struct_0(X29))|(~v5_pre_topc(X32,X30,X29)|~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X30),u1_struct_0(X29)))))))|(v2_struct_0(X31)|~m1_pre_topc(X31,X28))|(v2_struct_0(X30)|~m1_pre_topc(X30,X28))|(v2_struct_0(X29)|(~v2_pre_topc(X29)|~l1_pre_topc(X29)))|(v2_struct_0(X28)|(~v2_pre_topc(X28)|~l1_pre_topc(X28))))&((v5_pre_topc(k3_tmap_1(X28,X29,X30,X31,X32),X31,X29)|~m1_pre_topc(X31,X30)|(~v1_funct_1(X32)|(~v1_funct_2(X32,u1_struct_0(X30),u1_struct_0(X29))|(~v5_pre_topc(X32,X30,X29)|~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X30),u1_struct_0(X29)))))))|(v2_struct_0(X31)|~m1_pre_topc(X31,X28))|(v2_struct_0(X30)|~m1_pre_topc(X30,X28))|(v2_struct_0(X29)|(~v2_pre_topc(X29)|~l1_pre_topc(X29)))|(v2_struct_0(X28)|(~v2_pre_topc(X28)|~l1_pre_topc(X28))))&(m1_subset_1(k3_tmap_1(X28,X29,X30,X31,X32),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X31),u1_struct_0(X29))))|~m1_pre_topc(X31,X30)|(~v1_funct_1(X32)|(~v1_funct_2(X32,u1_struct_0(X30),u1_struct_0(X29))|(~v5_pre_topc(X32,X30,X29)|~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X30),u1_struct_0(X29)))))))|(v2_struct_0(X31)|~m1_pre_topc(X31,X28))|(v2_struct_0(X30)|~m1_pre_topc(X30,X28))|(v2_struct_0(X29)|(~v2_pre_topc(X29)|~l1_pre_topc(X29)))|(v2_struct_0(X28)|(~v2_pre_topc(X28)|~l1_pre_topc(X28))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t89_tmap_1])])])])])).
cnf(c_0_20, negated_conjecture, (r2_funct_2(u1_struct_0(X1),u1_struct_0(esk2_0),k3_tmap_1(X2,esk2_0,X3,X1,k3_tmap_1(X2,esk2_0,esk4_0,X3,esk6_0)),k3_tmap_1(X2,esk2_0,esk4_0,X1,esk6_0))|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X1)|~m1_pre_topc(X3,esk4_0)|~m1_pre_topc(esk4_0,X2)|~m1_pre_topc(X1,X3)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11]), c_0_12]), c_0_13]), c_0_14])]), c_0_15]), c_0_16])).
cnf(c_0_21, negated_conjecture, (m1_pre_topc(esk5_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_22, negated_conjecture, (m1_pre_topc(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_23, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_24, negated_conjecture, (~v2_struct_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_25, plain, (m1_subset_1(k3_tmap_1(X1,X2,X3,X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_26, negated_conjecture, (v1_funct_2(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0),u1_struct_0(esk3_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_27, negated_conjecture, (m1_subset_1(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_28, negated_conjecture, (v1_funct_1(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_29, negated_conjecture, (v1_funct_1(k3_tmap_1(X1,esk2_0,esk4_0,X2,esk6_0))|v2_struct_0(X1)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_10]), c_0_11]), c_0_12]), c_0_13]), c_0_14])]), c_0_16])).
cnf(c_0_30, negated_conjecture, (m1_pre_topc(esk5_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_31, negated_conjecture, (m1_pre_topc(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_32, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_33, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_34, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_35, negated_conjecture, (v1_funct_2(k3_tmap_1(X1,esk2_0,esk4_0,X2,esk6_0),u1_struct_0(X2),u1_struct_0(esk2_0))|v2_struct_0(X1)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_10]), c_0_11]), c_0_12]), c_0_13]), c_0_14])]), c_0_16])).
cnf(c_0_36, plain, (v5_pre_topc(k3_tmap_1(X1,X2,X3,X4,X5),X4,X2)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X4,X3)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v5_pre_topc(X5,X3,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_37, negated_conjecture, (v5_pre_topc(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0),esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_38, plain, ![X13, X14, X15, X16]:((~r2_funct_2(X13,X14,X15,X16)|X15=X16|(~v1_funct_1(X15)|(~v1_funct_2(X15,X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))|(~v1_funct_1(X16)|(~v1_funct_2(X16,X13,X14)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X13,X14)))))))&(X15!=X16|r2_funct_2(X13,X14,X15,X16)|(~v1_funct_1(X15)|(~v1_funct_2(X15,X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))|(~v1_funct_1(X16)|(~v1_funct_2(X16,X13,X14)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X13,X14)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r2_funct_2])])])).
cnf(c_0_39, negated_conjecture, (r2_funct_2(u1_struct_0(esk5_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,esk3_0,esk5_0,k3_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk6_0)),k3_tmap_1(X1,esk2_0,esk4_0,esk5_0,esk6_0))|v2_struct_0(X1)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(esk5_0,X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])]), c_0_23]), c_0_24])).
cnf(c_0_40, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_41, negated_conjecture, (m1_subset_1(k3_tmap_1(X1,esk2_0,esk4_0,X2,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|v2_struct_0(X1)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_10]), c_0_11]), c_0_12]), c_0_13]), c_0_14])]), c_0_16])).
cnf(c_0_42, negated_conjecture, (m1_subset_1(k3_tmap_1(X1,esk2_0,esk3_0,X2,k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_28]), c_0_13]), c_0_14])]), c_0_16])).
cnf(c_0_43, negated_conjecture, (v1_funct_2(k3_tmap_1(X1,esk2_0,esk3_0,X2,k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0)),u1_struct_0(X2),u1_struct_0(esk2_0))|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_26]), c_0_27]), c_0_28]), c_0_13]), c_0_14])]), c_0_16])).
cnf(c_0_44, negated_conjecture, (v1_funct_1(k3_tmap_1(X1,esk2_0,esk3_0,X2,k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0)))|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_26]), c_0_27]), c_0_28]), c_0_13]), c_0_14])]), c_0_16])).
cnf(c_0_45, negated_conjecture, (~v1_funct_1(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk5_0,esk6_0))|~v1_funct_2(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk5_0,esk6_0),u1_struct_0(esk5_0),u1_struct_0(esk2_0))|~v5_pre_topc(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk5_0,esk6_0),esk5_0,esk2_0)|~m1_subset_1(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_46, negated_conjecture, (v1_funct_1(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_32]), c_0_33])]), c_0_34])).
cnf(c_0_47, negated_conjecture, (v1_funct_2(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk5_0,esk6_0),u1_struct_0(esk5_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_30]), c_0_31]), c_0_32]), c_0_33])]), c_0_34])).
cnf(c_0_48, negated_conjecture, (v5_pre_topc(k3_tmap_1(X1,esk2_0,esk3_0,X2,k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0)),X2,esk2_0)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_pre_topc(X2,esk3_0)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_27]), c_0_26]), c_0_28]), c_0_13]), c_0_14])]), c_0_23]), c_0_16])).
cnf(c_0_49, plain, (X3=X4|~r2_funct_2(X1,X2,X3,X4)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X4)|~v1_funct_2(X4,X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, negated_conjecture, (r2_funct_2(u1_struct_0(esk5_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0,k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0)),k3_tmap_1(esk1_0,esk2_0,esk4_0,esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_31]), c_0_30]), c_0_40]), c_0_32]), c_0_33])]), c_0_34])).
cnf(c_0_51, negated_conjecture, (m1_subset_1(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(esk2_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_30]), c_0_31]), c_0_32]), c_0_33])]), c_0_34])).
cnf(c_0_52, negated_conjecture, (m1_subset_1(k3_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0,k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(esk2_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_30]), c_0_40]), c_0_32]), c_0_33])]), c_0_34])).
cnf(c_0_53, negated_conjecture, (v1_funct_2(k3_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0,k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0)),u1_struct_0(esk5_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_30]), c_0_40]), c_0_32]), c_0_33])]), c_0_34])).
cnf(c_0_54, negated_conjecture, (v1_funct_1(k3_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0,k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_30]), c_0_40]), c_0_32]), c_0_33])]), c_0_34])).
cnf(c_0_55, negated_conjecture, (~v5_pre_topc(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk5_0,esk6_0),esk5_0,esk2_0)|~m1_subset_1(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_46])]), c_0_47])])).
cnf(c_0_56, negated_conjecture, (v5_pre_topc(k3_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0,k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0)),esk5_0,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_30]), c_0_21]), c_0_40]), c_0_32]), c_0_33])]), c_0_34]), c_0_24])).
cnf(c_0_57, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0,k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0))=k3_tmap_1(esk1_0,esk2_0,esk4_0,esk5_0,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51]), c_0_52]), c_0_47]), c_0_53]), c_0_46]), c_0_54])])).
cnf(c_0_58, negated_conjecture, (~v5_pre_topc(k3_tmap_1(esk1_0,esk2_0,esk4_0,esk5_0,esk6_0),esk5_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_51])])).
cnf(c_0_59, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_57]), c_0_58]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 60
# Proof object clause steps            : 49
# Proof object formula steps           : 11
# Proof object conjectures             : 46
# Proof object clause conjectures      : 43
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 113
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 32
# Processed clauses                    : 167
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 164
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 19
# Generated clauses                    : 145
# ...of the previous two non-trivial   : 151
# Contextual simplify-reflections      : 0
# Paramodulations                      : 144
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 115
#    Positive orientable unit clauses  : 59
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 50
# Current number of unprocessed clauses: 43
# ...number of literals in the above   : 362
# Current number of archived formulas  : 0
# Current number of archived clauses   : 48
# Clause-clause subsumption calls (NU) : 1170
# Rec. Clause-clause subsumption calls : 75
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 305
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 238
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 13029

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.003 s
# Total time               : 0.015 s
# Maximum resident set size: 28792 pages
