# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_zfmisc_1, axiom, ![X1, X2, X3, X4]:k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4), file('data/MPTP/proofs/57178389865d4aa2b079593d7503c184.E_input', d4_zfmisc_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/57178389865d4aa2b079593d7503c184.E_input', d3_zfmisc_1)).
fof(t82_mcart_1, conjecture, ![X1, X2, X3, X4, X5, X6]:(m1_subset_1(X6,k4_zfmisc_1(X1,X2,X3,X4))=>(![X7]:(m1_subset_1(X7,X1)=>![X8]:(m1_subset_1(X8,X2)=>![X9]:(m1_subset_1(X9,X3)=>![X10]:(m1_subset_1(X10,X4)=>(X6=k4_mcart_1(X7,X8,X9,X10)=>X5=X10)))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|(X4=k1_xboole_0|X5=k11_mcart_1(X1,X2,X3,X4,X6))))))), file('data/MPTP/proofs/57178389865d4aa2b079593d7503c184.E_input', t82_mcart_1)).
fof(t81_mcart_1, axiom, ![X1, X2, X3, X4, X5, X6]:(m1_subset_1(X6,k4_zfmisc_1(X1,X2,X3,X4))=>(![X7]:(m1_subset_1(X7,X1)=>![X8]:(m1_subset_1(X8,X2)=>![X9]:(m1_subset_1(X9,X3)=>![X10]:(m1_subset_1(X10,X4)=>(X6=k4_mcart_1(X7,X8,X9,X10)=>X5=X9)))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|(X4=k1_xboole_0|X5=k10_mcart_1(X1,X2,X3,X4,X6))))))), file('data/MPTP/proofs/57178389865d4aa2b079593d7503c184.E_input', t81_mcart_1)).
fof(t78_mcart_1, axiom, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&?[X6, X7, X8, X9]:(X5=k4_mcart_1(X6,X7,X8,X9)&~((k8_mcart_1(X1,X2,X3,X4,X5)=X6&(k9_mcart_1(X1,X2,X3,X4,X5)=X7&(k10_mcart_1(X1,X2,X3,X4,X5)=X8&k11_mcart_1(X1,X2,X3,X4,X5)=X9))))))))))), file('data/MPTP/proofs/57178389865d4aa2b079593d7503c184.E_input', t78_mcart_1)).
fof(t61_mcart_1, axiom, ![X1, X2, X3, X4]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&~(![X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>(k8_mcart_1(X1,X2,X3,X4,X5)=k1_mcart_1(k1_mcart_1(k1_mcart_1(X5)))&(k9_mcart_1(X1,X2,X3,X4,X5)=k2_mcart_1(k1_mcart_1(k1_mcart_1(X5)))&(k10_mcart_1(X1,X2,X3,X4,X5)=k2_mcart_1(k1_mcart_1(X5))&k11_mcart_1(X1,X2,X3,X4,X5)=k2_mcart_1(X5))))))))))), file('data/MPTP/proofs/57178389865d4aa2b079593d7503c184.E_input', t61_mcart_1)).
fof(c_0_6, plain, ![X11, X12, X13, X14]:k4_zfmisc_1(X11,X12,X13,X14)=k2_zfmisc_1(k3_zfmisc_1(X11,X12,X13),X14), inference(variable_rename,[status(thm)],[d4_zfmisc_1])).
fof(c_0_7, plain, ![X15, X16, X17]:k3_zfmisc_1(X15,X16,X17)=k2_zfmisc_1(k2_zfmisc_1(X15,X16),X17), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_8, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6]:(m1_subset_1(X6,k4_zfmisc_1(X1,X2,X3,X4))=>(![X7]:(m1_subset_1(X7,X1)=>![X8]:(m1_subset_1(X8,X2)=>![X9]:(m1_subset_1(X9,X3)=>![X10]:(m1_subset_1(X10,X4)=>(X6=k4_mcart_1(X7,X8,X9,X10)=>X5=X10)))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|(X4=k1_xboole_0|X5=k11_mcart_1(X1,X2,X3,X4,X6)))))))), inference(assume_negation,[status(cth)],[t82_mcart_1])).
fof(c_0_9, plain, ![X28, X29, X30, X31, X32, X33]:((m1_subset_1(esk7_6(X28,X29,X30,X31,X32,X33),X28)|(X28=k1_xboole_0|(X29=k1_xboole_0|(X30=k1_xboole_0|(X31=k1_xboole_0|X32=k10_mcart_1(X28,X29,X30,X31,X33)))))|~m1_subset_1(X33,k4_zfmisc_1(X28,X29,X30,X31)))&((m1_subset_1(esk8_6(X28,X29,X30,X31,X32,X33),X29)|(X28=k1_xboole_0|(X29=k1_xboole_0|(X30=k1_xboole_0|(X31=k1_xboole_0|X32=k10_mcart_1(X28,X29,X30,X31,X33)))))|~m1_subset_1(X33,k4_zfmisc_1(X28,X29,X30,X31)))&((m1_subset_1(esk9_6(X28,X29,X30,X31,X32,X33),X30)|(X28=k1_xboole_0|(X29=k1_xboole_0|(X30=k1_xboole_0|(X31=k1_xboole_0|X32=k10_mcart_1(X28,X29,X30,X31,X33)))))|~m1_subset_1(X33,k4_zfmisc_1(X28,X29,X30,X31)))&((m1_subset_1(esk10_6(X28,X29,X30,X31,X32,X33),X31)|(X28=k1_xboole_0|(X29=k1_xboole_0|(X30=k1_xboole_0|(X31=k1_xboole_0|X32=k10_mcart_1(X28,X29,X30,X31,X33)))))|~m1_subset_1(X33,k4_zfmisc_1(X28,X29,X30,X31)))&((X33=k4_mcart_1(esk7_6(X28,X29,X30,X31,X32,X33),esk8_6(X28,X29,X30,X31,X32,X33),esk9_6(X28,X29,X30,X31,X32,X33),esk10_6(X28,X29,X30,X31,X32,X33))|(X28=k1_xboole_0|(X29=k1_xboole_0|(X30=k1_xboole_0|(X31=k1_xboole_0|X32=k10_mcart_1(X28,X29,X30,X31,X33)))))|~m1_subset_1(X33,k4_zfmisc_1(X28,X29,X30,X31)))&(X32!=esk9_6(X28,X29,X30,X31,X32,X33)|(X28=k1_xboole_0|(X29=k1_xboole_0|(X30=k1_xboole_0|(X31=k1_xboole_0|X32=k10_mcart_1(X28,X29,X30,X31,X33)))))|~m1_subset_1(X33,k4_zfmisc_1(X28,X29,X30,X31)))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t81_mcart_1])])])])).
cnf(c_0_10, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_12, negated_conjecture, ![X24, X25, X26, X27]:(m1_subset_1(esk6_0,k4_zfmisc_1(esk1_0,esk2_0,esk3_0,esk4_0))&((~m1_subset_1(X24,esk1_0)|(~m1_subset_1(X25,esk2_0)|(~m1_subset_1(X26,esk3_0)|(~m1_subset_1(X27,esk4_0)|(esk6_0!=k4_mcart_1(X24,X25,X26,X27)|esk5_0=X27)))))&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&(esk4_0!=k1_xboole_0&esk5_0!=k11_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
cnf(c_0_13, plain, (X1=k4_mcart_1(esk7_6(X2,X3,X4,X5,X6,X1),esk8_6(X2,X3,X4,X5,X6,X1),esk9_6(X2,X3,X4,X5,X6,X1),esk10_6(X2,X3,X4,X5,X6,X1))|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|X5=k1_xboole_0|X6=k10_mcart_1(X2,X3,X4,X5,X1)|~m1_subset_1(X1,k4_zfmisc_1(X2,X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15, negated_conjecture, (esk5_0=X4|~m1_subset_1(X1,esk1_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X4,esk4_0)|esk6_0!=k4_mcart_1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, plain, (X5=k1_xboole_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X6=k10_mcart_1(X2,X3,X4,X5,X1)|X1=k4_mcart_1(esk7_6(X2,X3,X4,X5,X6,X1),esk8_6(X2,X3,X4,X5,X6,X1),esk9_6(X2,X3,X4,X5,X6,X1),esk10_6(X2,X3,X4,X5,X6,X1))|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4),X5))), inference(rw,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_17, plain, (m1_subset_1(esk10_6(X1,X2,X3,X4,X5,X6),X4)|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|X5=k10_mcart_1(X1,X2,X3,X4,X6)|~m1_subset_1(X6,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_18, negated_conjecture, (esk10_6(X1,X2,X3,X4,X5,esk6_0)=esk5_0|X5=k10_mcart_1(X1,X2,X3,X4,esk6_0)|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|~m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))|~m1_subset_1(esk10_6(X1,X2,X3,X4,X5,esk6_0),esk4_0)|~m1_subset_1(esk9_6(X1,X2,X3,X4,X5,esk6_0),esk3_0)|~m1_subset_1(esk8_6(X1,X2,X3,X4,X5,esk6_0),esk2_0)|~m1_subset_1(esk7_6(X1,X2,X3,X4,X5,esk6_0),esk1_0)), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16])])).
cnf(c_0_19, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X5=k10_mcart_1(X1,X2,X3,X4,X6)|m1_subset_1(esk10_6(X1,X2,X3,X4,X5,X6),X4)|~m1_subset_1(X6,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[c_0_17, c_0_14])).
cnf(c_0_20, negated_conjecture, (esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, plain, (m1_subset_1(esk8_6(X1,X2,X3,X4,X5,X6),X2)|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|X5=k10_mcart_1(X1,X2,X3,X4,X6)|~m1_subset_1(X6,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_22, negated_conjecture, (esk10_6(X1,X2,X3,esk4_0,X4,esk6_0)=esk5_0|X4=k10_mcart_1(X1,X2,X3,esk4_0,esk6_0)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),esk4_0))|~m1_subset_1(esk9_6(X1,X2,X3,esk4_0,X4,esk6_0),esk3_0)|~m1_subset_1(esk8_6(X1,X2,X3,esk4_0,X4,esk6_0),esk2_0)|~m1_subset_1(esk7_6(X1,X2,X3,esk4_0,X4,esk6_0),esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])).
cnf(c_0_23, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X5=k10_mcart_1(X1,X2,X3,X4,X6)|m1_subset_1(esk8_6(X1,X2,X3,X4,X5,X6),X2)|~m1_subset_1(X6,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[c_0_21, c_0_14])).
cnf(c_0_24, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_25, plain, (m1_subset_1(esk7_6(X1,X2,X3,X4,X5,X6),X1)|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|X5=k10_mcart_1(X1,X2,X3,X4,X6)|~m1_subset_1(X6,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_26, plain, ![X38, X39, X40, X41, X42, X43, X44, X45, X46]:((k8_mcart_1(X38,X39,X40,X41,X42)=X43|X42!=k4_mcart_1(X43,X44,X45,X46)|X41=k1_xboole_0|X40=k1_xboole_0|X39=k1_xboole_0|X38=k1_xboole_0|~m1_subset_1(X42,k4_zfmisc_1(X38,X39,X40,X41)))&((k9_mcart_1(X38,X39,X40,X41,X42)=X44|X42!=k4_mcart_1(X43,X44,X45,X46)|X41=k1_xboole_0|X40=k1_xboole_0|X39=k1_xboole_0|X38=k1_xboole_0|~m1_subset_1(X42,k4_zfmisc_1(X38,X39,X40,X41)))&((k10_mcart_1(X38,X39,X40,X41,X42)=X45|X42!=k4_mcart_1(X43,X44,X45,X46)|X41=k1_xboole_0|X40=k1_xboole_0|X39=k1_xboole_0|X38=k1_xboole_0|~m1_subset_1(X42,k4_zfmisc_1(X38,X39,X40,X41)))&(k11_mcart_1(X38,X39,X40,X41,X42)=X46|X42!=k4_mcart_1(X43,X44,X45,X46)|X41=k1_xboole_0|X40=k1_xboole_0|X39=k1_xboole_0|X38=k1_xboole_0|~m1_subset_1(X42,k4_zfmisc_1(X38,X39,X40,X41)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t78_mcart_1])])])])])).
cnf(c_0_27, negated_conjecture, (esk10_6(X1,esk2_0,X2,esk4_0,X3,esk6_0)=esk5_0|X3=k10_mcart_1(X1,esk2_0,X2,esk4_0,esk6_0)|X1=k1_xboole_0|X2=k1_xboole_0|~m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,esk2_0),X2),esk4_0))|~m1_subset_1(esk9_6(X1,esk2_0,X2,esk4_0,X3,esk6_0),esk3_0)|~m1_subset_1(esk7_6(X1,esk2_0,X2,esk4_0,X3,esk6_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_20])).
cnf(c_0_28, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X5=k10_mcart_1(X1,X2,X3,X4,X6)|m1_subset_1(esk7_6(X1,X2,X3,X4,X5,X6),X1)|~m1_subset_1(X6,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[c_0_25, c_0_14])).
cnf(c_0_29, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_30, plain, (m1_subset_1(esk9_6(X1,X2,X3,X4,X5,X6),X3)|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|X5=k10_mcart_1(X1,X2,X3,X4,X6)|~m1_subset_1(X6,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk6_0,k4_zfmisc_1(esk1_0,esk2_0,esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_32, plain, ![X47, X48, X49, X50, X51]:((k8_mcart_1(X47,X48,X49,X50,X51)=k1_mcart_1(k1_mcart_1(k1_mcart_1(X51)))|~m1_subset_1(X51,k4_zfmisc_1(X47,X48,X49,X50))|X50=k1_xboole_0|X49=k1_xboole_0|X48=k1_xboole_0|X47=k1_xboole_0)&((k9_mcart_1(X47,X48,X49,X50,X51)=k2_mcart_1(k1_mcart_1(k1_mcart_1(X51)))|~m1_subset_1(X51,k4_zfmisc_1(X47,X48,X49,X50))|X50=k1_xboole_0|X49=k1_xboole_0|X48=k1_xboole_0|X47=k1_xboole_0)&((k10_mcart_1(X47,X48,X49,X50,X51)=k2_mcart_1(k1_mcart_1(X51))|~m1_subset_1(X51,k4_zfmisc_1(X47,X48,X49,X50))|X50=k1_xboole_0|X49=k1_xboole_0|X48=k1_xboole_0|X47=k1_xboole_0)&(k11_mcart_1(X47,X48,X49,X50,X51)=k2_mcart_1(X51)|~m1_subset_1(X51,k4_zfmisc_1(X47,X48,X49,X50))|X50=k1_xboole_0|X49=k1_xboole_0|X48=k1_xboole_0|X47=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t61_mcart_1])])])])])).
cnf(c_0_33, plain, (k11_mcart_1(X1,X2,X3,X4,X5)=X6|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X5!=k4_mcart_1(X7,X8,X9,X6)|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, negated_conjecture, (esk10_6(esk1_0,esk2_0,X1,esk4_0,X2,esk6_0)=esk5_0|X2=k10_mcart_1(esk1_0,esk2_0,X1,esk4_0,esk6_0)|X1=k1_xboole_0|~m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),X1),esk4_0))|~m1_subset_1(esk9_6(esk1_0,esk2_0,X1,esk4_0,X2,esk6_0),esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_20]), c_0_24])).
cnf(c_0_35, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X5=k10_mcart_1(X1,X2,X3,X4,X6)|m1_subset_1(esk9_6(X1,X2,X3,X4,X5,X6),X3)|~m1_subset_1(X6,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[c_0_30, c_0_14])).
cnf(c_0_36, negated_conjecture, (m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),esk4_0))), inference(rw,[status(thm)],[c_0_31, c_0_14])).
cnf(c_0_37, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_38, plain, (k11_mcart_1(X1,X2,X3,X4,X5)=k2_mcart_1(X5)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k11_mcart_1(X1,X2,X3,X4,X5)=X6|X5!=k4_mcart_1(X7,X8,X9,X6)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[c_0_33, c_0_14])).
cnf(c_0_40, negated_conjecture, (esk10_6(esk1_0,esk2_0,esk3_0,esk4_0,X1,esk6_0)=esk5_0|X1=k10_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])]), c_0_37]), c_0_20]), c_0_24]), c_0_29])).
cnf(c_0_41, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k11_mcart_1(X1,X2,X3,X4,X5)=k2_mcart_1(X5)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[c_0_38, c_0_14])).
cnf(c_0_42, plain, (k11_mcart_1(X1,X2,X3,X4,k4_mcart_1(X5,X6,X7,X8))=X8|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(k4_mcart_1(X5,X6,X7,X8),k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(er,[status(thm)],[c_0_39])).
cnf(c_0_43, negated_conjecture, (k4_mcart_1(esk7_6(esk1_0,esk2_0,esk3_0,esk4_0,X1,esk6_0),esk8_6(esk1_0,esk2_0,esk3_0,esk4_0,X1,esk6_0),esk9_6(esk1_0,esk2_0,esk3_0,esk4_0,X1,esk6_0),esk5_0)=esk6_0|X1=k10_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_40]), c_0_36])]), c_0_20]), c_0_37]), c_0_24]), c_0_29])).
cnf(c_0_44, negated_conjecture, (esk5_0!=k11_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_45, negated_conjecture, (k11_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)=k2_mcart_1(esk6_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_36]), c_0_20]), c_0_37]), c_0_24]), c_0_29])).
cnf(c_0_46, negated_conjecture, (X1=k10_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)|k11_mcart_1(X2,X3,X4,X5,esk6_0)=esk5_0|X5=k1_xboole_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|~m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4),X5))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_47, negated_conjecture, (k2_mcart_1(esk6_0)!=esk5_0), inference(rw,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_48, negated_conjecture, (X1=k10_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_36]), c_0_45]), c_0_47]), c_0_20]), c_0_37]), c_0_24]), c_0_29])).
cnf(c_0_49, negated_conjecture, (X1=X2), inference(spm,[status(thm)],[c_0_48, c_0_48])).
cnf(c_0_50, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_29, c_0_49]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 51
# Proof object clause steps            : 38
# Proof object formula steps           : 13
# Proof object conjectures             : 23
# Proof object clause conjectures      : 20
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 41
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 21
# Processed clauses                    : 67
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 67
# Other redundant clauses eliminated   : 19
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 30
# Generated clauses                    : 361
# ...of the previous two non-trivial   : 359
# Contextual simplify-reflections      : 0
# Paramodulations                      : 337
# Factorizations                       : 0
# Equation resolutions                 : 19
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2
#    Positive orientable unit clauses  : 1
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 80
# ...number of literals in the above   : 734
# Current number of archived formulas  : 0
# Current number of archived clauses   : 63
# Clause-clause subsumption calls (NU) : 146
# Rec. Clause-clause subsumption calls : 16
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 30
# Rewrite failures with RHS unbound    : 29
# BW rewrite match attempts            : 113
# BW rewrite match successes           : 110
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8003

# -------------------------------------------------
# User time                : 0.014 s
# System time              : 0.000 s
# Total time               : 0.014 s
# Maximum resident set size: 28660 pages
