# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/578b41da5cc646de90780f1c17a45d1a.E_input', d2_zfmisc_1)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/578b41da5cc646de90780f1c17a45d1a.E_input', t7_mcart_1)).
fof(t73_mcart_1, conjecture, ![X1, X2, X3, X4, X5, X6]:(r2_hidden(k3_mcart_1(X1,X2,X3),k3_zfmisc_1(X4,X5,X6))<=>(r2_hidden(X1,X4)&(r2_hidden(X2,X5)&r2_hidden(X3,X6)))), file('data/MPTP/proofs/578b41da5cc646de90780f1c17a45d1a.E_input', t73_mcart_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/578b41da5cc646de90780f1c17a45d1a.E_input', d3_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/578b41da5cc646de90780f1c17a45d1a.E_input', d3_zfmisc_1)).
fof(c_0_5, plain, ![X7, X8, X9, X10, X13, X14, X15, X16, X17, X18, X20, X21]:((((r2_hidden(esk1_4(X7,X8,X9,X10),X7)|~r2_hidden(X10,X9)|X9!=k2_zfmisc_1(X7,X8))&((r2_hidden(esk2_4(X7,X8,X9,X10),X8)|~r2_hidden(X10,X9)|X9!=k2_zfmisc_1(X7,X8))&(X10=k4_tarski(esk1_4(X7,X8,X9,X10),esk2_4(X7,X8,X9,X10))|~r2_hidden(X10,X9)|X9!=k2_zfmisc_1(X7,X8))))&(~r2_hidden(X14,X7)|(~r2_hidden(X15,X8)|X13!=k4_tarski(X14,X15))|r2_hidden(X13,X9)|X9!=k2_zfmisc_1(X7,X8)))&((~r2_hidden(esk3_3(X16,X17,X18),X18)|(~r2_hidden(X20,X16)|(~r2_hidden(X21,X17)|esk3_3(X16,X17,X18)!=k4_tarski(X20,X21)))|X18=k2_zfmisc_1(X16,X17))&((r2_hidden(esk4_3(X16,X17,X18),X16)|r2_hidden(esk3_3(X16,X17,X18),X18)|X18=k2_zfmisc_1(X16,X17))&((r2_hidden(esk5_3(X16,X17,X18),X17)|r2_hidden(esk3_3(X16,X17,X18),X18)|X18=k2_zfmisc_1(X16,X17))&(esk3_3(X16,X17,X18)=k4_tarski(esk4_3(X16,X17,X18),esk5_3(X16,X17,X18))|r2_hidden(esk3_3(X16,X17,X18),X18)|X18=k2_zfmisc_1(X16,X17)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
fof(c_0_6, plain, ![X24, X25]:(k1_mcart_1(k4_tarski(X24,X25))=X24&k2_mcart_1(k4_tarski(X24,X25))=X25), inference(variable_rename,[status(thm)],[t7_mcart_1])).
cnf(c_0_7, plain, (X1=k4_tarski(esk1_4(X2,X3,X4,X1),esk2_4(X2,X3,X4,X1))|~r2_hidden(X1,X4)|X4!=k2_zfmisc_1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_8, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6]:(r2_hidden(k3_mcart_1(X1,X2,X3),k3_zfmisc_1(X4,X5,X6))<=>(r2_hidden(X1,X4)&(r2_hidden(X2,X5)&r2_hidden(X3,X6))))), inference(assume_negation,[status(cth)],[t73_mcart_1])).
cnf(c_0_9, plain, (r2_hidden(esk1_4(X1,X2,X3,X4),X1)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11, plain, (k4_tarski(esk1_4(X1,X2,k2_zfmisc_1(X1,X2),X3),esk2_4(X1,X2,k2_zfmisc_1(X1,X2),X3))=X3|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_7])).
fof(c_0_12, negated_conjecture, ((~r2_hidden(k3_mcart_1(esk6_0,esk7_0,esk8_0),k3_zfmisc_1(esk9_0,esk10_0,esk11_0))|(~r2_hidden(esk6_0,esk9_0)|(~r2_hidden(esk7_0,esk10_0)|~r2_hidden(esk8_0,esk11_0))))&((r2_hidden(esk6_0,esk9_0)|r2_hidden(k3_mcart_1(esk6_0,esk7_0,esk8_0),k3_zfmisc_1(esk9_0,esk10_0,esk11_0)))&((r2_hidden(esk7_0,esk10_0)|r2_hidden(k3_mcart_1(esk6_0,esk7_0,esk8_0),k3_zfmisc_1(esk9_0,esk10_0,esk11_0)))&(r2_hidden(esk8_0,esk11_0)|r2_hidden(k3_mcart_1(esk6_0,esk7_0,esk8_0),k3_zfmisc_1(esk9_0,esk10_0,esk11_0)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
fof(c_0_13, plain, ![X32, X33, X34]:k3_mcart_1(X32,X33,X34)=k4_tarski(k4_tarski(X32,X33),X34), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_14, plain, ![X35, X36, X37]:k3_zfmisc_1(X35,X36,X37)=k2_zfmisc_1(k2_zfmisc_1(X35,X36),X37), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
cnf(c_0_15, plain, (r2_hidden(esk1_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X1)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_9])).
cnf(c_0_16, plain, (esk1_4(X1,X2,k2_zfmisc_1(X1,X2),X3)=k1_mcart_1(X3)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_17, negated_conjecture, (r2_hidden(esk6_0,esk9_0)|r2_hidden(k3_mcart_1(esk6_0,esk7_0,esk8_0),k3_zfmisc_1(esk9_0,esk10_0,esk11_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, negated_conjecture, (~r2_hidden(k3_mcart_1(esk6_0,esk7_0,esk8_0),k3_zfmisc_1(esk9_0,esk10_0,esk11_0))|~r2_hidden(esk6_0,esk9_0)|~r2_hidden(esk7_0,esk10_0)|~r2_hidden(esk8_0,esk11_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, plain, (r2_hidden(X5,X6)|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)|X5!=k4_tarski(X1,X3)|X6!=k2_zfmisc_1(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_22, plain, (r2_hidden(k1_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_23, negated_conjecture, (r2_hidden(esk6_0,esk9_0)|r2_hidden(k4_tarski(k4_tarski(esk6_0,esk7_0),esk8_0),k2_zfmisc_1(k2_zfmisc_1(esk9_0,esk10_0),esk11_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_18]), c_0_19])).
cnf(c_0_24, plain, (r2_hidden(esk2_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_25, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_26, negated_conjecture, (r2_hidden(esk7_0,esk10_0)|r2_hidden(k3_mcart_1(esk6_0,esk7_0,esk8_0),k3_zfmisc_1(esk9_0,esk10_0,esk11_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_27, negated_conjecture, (~r2_hidden(esk6_0,esk9_0)|~r2_hidden(esk7_0,esk10_0)|~r2_hidden(esk8_0,esk11_0)|~r2_hidden(k4_tarski(k4_tarski(esk6_0,esk7_0),esk8_0),k2_zfmisc_1(k2_zfmisc_1(esk9_0,esk10_0),esk11_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_18]), c_0_19])).
cnf(c_0_28, plain, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_21])])).
cnf(c_0_29, negated_conjecture, (r2_hidden(k4_tarski(esk6_0,esk7_0),k2_zfmisc_1(esk9_0,esk10_0))|r2_hidden(esk6_0,esk9_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_10])).
cnf(c_0_30, plain, (r2_hidden(esk2_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X2)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_24])).
cnf(c_0_31, plain, (esk2_4(X1,X2,k2_zfmisc_1(X1,X2),X3)=k2_mcart_1(X3)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_25, c_0_11])).
cnf(c_0_32, negated_conjecture, (r2_hidden(esk7_0,esk10_0)|r2_hidden(k4_tarski(k4_tarski(esk6_0,esk7_0),esk8_0),k2_zfmisc_1(k2_zfmisc_1(esk9_0,esk10_0),esk11_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_18]), c_0_19])).
cnf(c_0_33, negated_conjecture, (~r2_hidden(esk6_0,esk9_0)|~r2_hidden(esk7_0,esk10_0)|~r2_hidden(esk8_0,esk11_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_28])).
cnf(c_0_34, negated_conjecture, (r2_hidden(esk6_0,esk9_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_29]), c_0_10])])).
cnf(c_0_35, plain, (r2_hidden(k2_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X3,X2))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_36, negated_conjecture, (r2_hidden(k4_tarski(esk6_0,esk7_0),k2_zfmisc_1(esk9_0,esk10_0))|r2_hidden(esk7_0,esk10_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_32]), c_0_10])).
cnf(c_0_37, negated_conjecture, (r2_hidden(esk8_0,esk11_0)|r2_hidden(k3_mcart_1(esk6_0,esk7_0,esk8_0),k3_zfmisc_1(esk9_0,esk10_0,esk11_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_38, negated_conjecture, (~r2_hidden(esk7_0,esk10_0)|~r2_hidden(esk8_0,esk11_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_34])])).
cnf(c_0_39, negated_conjecture, (r2_hidden(esk7_0,esk10_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_25])])).
cnf(c_0_40, negated_conjecture, (r2_hidden(esk8_0,esk11_0)|r2_hidden(k4_tarski(k4_tarski(esk6_0,esk7_0),esk8_0),k2_zfmisc_1(k2_zfmisc_1(esk9_0,esk10_0),esk11_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_18]), c_0_19])).
cnf(c_0_41, negated_conjecture, (~r2_hidden(esk8_0,esk11_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_39])])).
cnf(c_0_42, negated_conjecture, (r2_hidden(k4_tarski(k4_tarski(esk6_0,esk7_0),esk8_0),k2_zfmisc_1(k2_zfmisc_1(esk9_0,esk10_0),esk11_0))), inference(sr,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_43, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_42]), c_0_25]), c_0_41]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 44
# Proof object clause steps            : 33
# Proof object formula steps           : 11
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 27
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 14
# Processed clauses                    : 45
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 45
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 6
# Generated clauses                    : 56
# ...of the previous two non-trivial   : 53
# Contextual simplify-reflections      : 1
# Paramodulations                      : 51
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 19
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 13
# Current number of unprocessed clauses: 34
# ...number of literals in the above   : 124
# Current number of archived formulas  : 0
# Current number of archived clauses   : 24
# Clause-clause subsumption calls (NU) : 38
# Rec. Clause-clause subsumption calls : 27
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 13
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2029

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28528 pages
