# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/57942b1fd1a34343943c2ed93c8cd3a1.E_input', d1_xboole_0)).
fof(t6_xboole_0, axiom, ![X1, X2]:~((r2_xboole_0(X1,X2)&![X3]:~((r2_hidden(X3,X2)&~(r2_hidden(X3,X1)))))), file('data/MPTP/proofs/57942b1fd1a34343943c2ed93c8cd3a1.E_input', t6_xboole_0)).
fof(d8_xboole_0, axiom, ![X1, X2]:(r2_xboole_0(X1,X2)<=>(r1_tarski(X1,X2)&~(X1=X2))), file('data/MPTP/proofs/57942b1fd1a34343943c2ed93c8cd3a1.E_input', d8_xboole_0)).
fof(t6_relset_1, conjecture, ![X1, X2, X3, X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=>~((r2_hidden(X1,X4)&![X5, X6]:~((X1=k4_tarski(X5,X6)&(r2_hidden(X5,X2)&r2_hidden(X6,X3))))))), file('data/MPTP/proofs/57942b1fd1a34343943c2ed93c8cd3a1.E_input', t6_relset_1)).
fof(d18_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(v4_relat_1(X2,X1)<=>r1_tarski(k1_relat_1(X2),X1))), file('data/MPTP/proofs/57942b1fd1a34343943c2ed93c8cd3a1.E_input', d18_relat_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('data/MPTP/proofs/57942b1fd1a34343943c2ed93c8cd3a1.E_input', cc2_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/57942b1fd1a34343943c2ed93c8cd3a1.E_input', cc1_relset_1)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('data/MPTP/proofs/57942b1fd1a34343943c2ed93c8cd3a1.E_input', cc1_subset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/57942b1fd1a34343943c2ed93c8cd3a1.E_input', t3_subset)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/57942b1fd1a34343943c2ed93c8cd3a1.E_input', t7_xboole_0)).
fof(t65_subset_1, axiom, ![X1, X2, X3, X4]:~((r2_hidden(X4,X3)&(r1_tarski(X3,k2_zfmisc_1(X1,X2))&![X5]:(m1_subset_1(X5,X1)=>![X6]:(m1_subset_1(X6,X2)=>~(X4=k4_tarski(X5,X6))))))), file('data/MPTP/proofs/57942b1fd1a34343943c2ed93c8cd3a1.E_input', t65_subset_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/57942b1fd1a34343943c2ed93c8cd3a1.E_input', t2_subset)).
fof(t21_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>r1_tarski(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))), file('data/MPTP/proofs/57942b1fd1a34343943c2ed93c8cd3a1.E_input', t21_relat_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/57942b1fd1a34343943c2ed93c8cd3a1.E_input', t113_zfmisc_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/57942b1fd1a34343943c2ed93c8cd3a1.E_input', fc1_xboole_0)).
fof(c_0_15, plain, ![X32, X33, X34]:((~v1_xboole_0(X32)|~r2_hidden(X33,X32))&(r2_hidden(esk6_1(X34),X34)|v1_xboole_0(X34))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_16, plain, ![X25, X26]:((r2_hidden(esk5_2(X25,X26),X26)|~r2_xboole_0(X25,X26))&(~r2_hidden(esk5_2(X25,X26),X25)|~r2_xboole_0(X25,X26))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_xboole_0])])])])])).
cnf(c_0_17, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_18, plain, (r2_hidden(esk5_2(X1,X2),X2)|~r2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_19, plain, ![X28, X29]:(((r1_tarski(X28,X29)|~r2_xboole_0(X28,X29))&(X28!=X29|~r2_xboole_0(X28,X29)))&(~r1_tarski(X28,X29)|X28=X29|r2_xboole_0(X28,X29))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_xboole_0])])])])).
fof(c_0_20, negated_conjecture, ~(![X1, X2, X3, X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=>~((r2_hidden(X1,X4)&![X5, X6]:~((X1=k4_tarski(X5,X6)&(r2_hidden(X5,X2)&r2_hidden(X6,X3)))))))), inference(assume_negation,[status(cth)],[t6_relset_1])).
cnf(c_0_21, plain, (~r2_xboole_0(X1,X2)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_22, plain, (X1=X2|r2_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_23, plain, ![X17, X18]:((~v4_relat_1(X18,X17)|r1_tarski(k1_relat_1(X18),X17)|~v1_relat_1(X18))&(~r1_tarski(k1_relat_1(X18),X17)|v4_relat_1(X18,X17)|~v1_relat_1(X18))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d18_relat_1])])])).
fof(c_0_24, plain, ![X19, X20, X21]:((v4_relat_1(X21,X19)|~m1_subset_1(X21,k1_zfmisc_1(k2_zfmisc_1(X19,X20))))&(v5_relat_1(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(k2_zfmisc_1(X19,X20))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_25, negated_conjecture, ![X11, X12]:(m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))&(r2_hidden(esk1_0,esk4_0)&(esk1_0!=k4_tarski(X11,X12)|(~r2_hidden(X11,esk2_0)|~r2_hidden(X12,esk3_0))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])])).
fof(c_0_26, plain, ![X22, X23, X24]:(~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(X22,X23)))|v1_relat_1(X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_27, plain, ![X15, X16]:(~v1_xboole_0(X15)|(~m1_subset_1(X16,k1_zfmisc_1(X15))|v1_xboole_0(X16))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
fof(c_0_28, plain, ![X45, X46]:((~m1_subset_1(X45,k1_zfmisc_1(X46))|r1_tarski(X45,X46))&(~r1_tarski(X45,X46)|m1_subset_1(X45,k1_zfmisc_1(X46)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_29, plain, (X1=X2|~r1_tarski(X1,X2)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_30, plain, (r1_tarski(k1_relat_1(X1),X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36, plain, (k1_relat_1(X1)=X2|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_37, negated_conjecture, (v4_relat_1(esk4_0,esk2_0)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_38, negated_conjecture, (v1_relat_1(esk4_0)), inference(spm,[status(thm)],[c_0_33, c_0_32])).
cnf(c_0_39, plain, (v1_xboole_0(X1)|~r1_tarski(X1,X2)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40, plain, (v4_relat_1(X1,X2)|~r1_tarski(k1_relat_1(X1),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_41, negated_conjecture, (k1_relat_1(esk4_0)=esk2_0|~v1_xboole_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])])).
fof(c_0_42, plain, ![X43]:(X43=k1_xboole_0|r2_hidden(esk9_1(X43),X43)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_43, plain, (v1_xboole_0(k1_relat_1(X1))|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_39, c_0_30])).
cnf(c_0_44, negated_conjecture, (v4_relat_1(esk4_0,X1)|~r1_tarski(esk2_0,X1)|~v1_xboole_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_38])])).
cnf(c_0_45, negated_conjecture, (r1_tarski(esk2_0,X1)|~v4_relat_1(esk4_0,X1)|~v1_xboole_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_41]), c_0_38])])).
fof(c_0_46, plain, ![X37, X38, X39, X40]:((m1_subset_1(esk7_4(X37,X38,X39,X40),X37)|~r1_tarski(X39,k2_zfmisc_1(X37,X38))|~r2_hidden(X40,X39))&((m1_subset_1(esk8_4(X37,X38,X39,X40),X38)|~r1_tarski(X39,k2_zfmisc_1(X37,X38))|~r2_hidden(X40,X39))&(X40=k4_tarski(esk7_4(X37,X38,X39,X40),esk8_4(X37,X38,X39,X40))|~r1_tarski(X39,k2_zfmisc_1(X37,X38))|~r2_hidden(X40,X39)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t65_subset_1])])])])])).
fof(c_0_47, plain, ![X13, X14]:(~m1_subset_1(X13,X14)|(v1_xboole_0(X14)|r2_hidden(X13,X14))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_48, plain, ![X36]:(~v1_relat_1(X36)|r1_tarski(X36,k2_zfmisc_1(k1_relat_1(X36),k2_relat_1(X36)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t21_relat_1])])).
cnf(c_0_49, plain, (X1=k1_xboole_0|r2_hidden(esk9_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_50, negated_conjecture, (v1_xboole_0(k1_relat_1(esk4_0))|~r1_tarski(esk2_0,X1)|~v1_xboole_0(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_38])]), c_0_39])).
cnf(c_0_51, negated_conjecture, (r1_tarski(esk2_0,esk2_0)|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_45, c_0_37])).
fof(c_0_52, plain, ![X30, X31]:((k2_zfmisc_1(X30,X31)!=k1_xboole_0|(X30=k1_xboole_0|X31=k1_xboole_0))&((X30!=k1_xboole_0|k2_zfmisc_1(X30,X31)=k1_xboole_0)&(X31!=k1_xboole_0|k2_zfmisc_1(X30,X31)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_53, negated_conjecture, (esk1_0!=k4_tarski(X1,X2)|~r2_hidden(X1,esk2_0)|~r2_hidden(X2,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_54, plain, (X1=k4_tarski(esk7_4(X2,X3,X4,X1),esk8_4(X2,X3,X4,X1))|~r1_tarski(X4,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,X4)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_55, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_56, plain, (m1_subset_1(esk8_4(X1,X2,X3,X4),X2)|~r1_tarski(X3,k2_zfmisc_1(X1,X2))|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_57, plain, (r1_tarski(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_58, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_17, c_0_49])).
cnf(c_0_59, negated_conjecture, (v1_xboole_0(k1_relat_1(esk4_0))|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_60, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_61, negated_conjecture, (r2_hidden(esk1_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_62, negated_conjecture, (~r1_tarski(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(esk8_4(X2,X3,X1,esk1_0),esk3_0)|~r2_hidden(esk7_4(X2,X3,X1,esk1_0),esk2_0)|~r2_hidden(esk1_0,X1)), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54])])).
cnf(c_0_63, plain, (v1_xboole_0(X1)|r2_hidden(esk8_4(X2,X1,X3,X4),X1)|~r1_tarski(X3,k2_zfmisc_1(X2,X1))|~r2_hidden(X4,X3)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_64, plain, (m1_subset_1(esk7_4(X1,X2,X3,X4),X1)|~r1_tarski(X3,k2_zfmisc_1(X1,X2))|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_65, plain, (v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))), inference(spm,[status(thm)],[c_0_39, c_0_57])).
cnf(c_0_66, negated_conjecture, (k1_relat_1(esk4_0)=k1_xboole_0|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_67, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_60])).
cnf(c_0_68, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_69, negated_conjecture, (~v1_xboole_0(esk4_0)), inference(spm,[status(thm)],[c_0_17, c_0_61])).
cnf(c_0_70, negated_conjecture, (v1_xboole_0(esk3_0)|~r1_tarski(X1,k2_zfmisc_1(X2,esk3_0))|~r2_hidden(esk7_4(X2,esk3_0,X1,esk1_0),esk2_0)|~r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_71, plain, (v1_xboole_0(X1)|r2_hidden(esk7_4(X1,X2,X3,X4),X1)|~r1_tarski(X3,k2_zfmisc_1(X1,X2))|~r2_hidden(X4,X3)), inference(spm,[status(thm)],[c_0_55, c_0_64])).
cnf(c_0_72, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_38]), c_0_67]), c_0_68])]), c_0_69])).
cnf(c_0_73, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_74, negated_conjecture, (v1_xboole_0(esk3_0)|~r1_tarski(X1,k2_zfmisc_1(esk2_0,esk3_0))|~r2_hidden(esk1_0,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_72])).
cnf(c_0_75, negated_conjecture, (r1_tarski(esk4_0,k2_zfmisc_1(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_73, c_0_32])).
cnf(c_0_76, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_77, negated_conjecture, (v1_xboole_0(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_61])])).
cnf(c_0_78, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_76])).
cnf(c_0_79, negated_conjecture, (k1_xboole_0=esk3_0), inference(spm,[status(thm)],[c_0_58, c_0_77])).
cnf(c_0_80, negated_conjecture, (~v1_xboole_0(k2_zfmisc_1(esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_32]), c_0_69])).
cnf(c_0_81, plain, (k2_zfmisc_1(X1,esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_79]), c_0_79])).
cnf(c_0_82, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_80, c_0_81]), c_0_77])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 83
# Proof object clause steps            : 53
# Proof object formula steps           : 30
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 26
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 28
# Processed clauses                    : 255
# ...of these trivial                  : 0
# ...subsumed                          : 98
# ...remaining for further processing  : 157
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 11
# Backward-rewritten                   : 62
# Generated clauses                    : 399
# ...of the previous two non-trivial   : 398
# Contextual simplify-reflections      : 17
# Paramodulations                      : 395
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 81
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 70
# Current number of unprocessed clauses: 167
# ...number of literals in the above   : 578
# Current number of archived formulas  : 0
# Current number of archived clauses   : 73
# Clause-clause subsumption calls (NU) : 3105
# Rec. Clause-clause subsumption calls : 2216
# Non-unit clause-clause subsumptions  : 84
# Unit Clause-clause subsumption calls : 189
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6858

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.000 s
# Total time               : 0.015 s
# Maximum resident set size: 28716 pages
