# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/579d6d627e734c97be966e07f621e954.E_input', d5_xboole_0)).
fof(t4_boole, axiom, ![X1]:k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0, file('data/MPTP/proofs/579d6d627e734c97be966e07f621e954.E_input', t4_boole)).
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/579d6d627e734c97be966e07f621e954.E_input', d2_zfmisc_1)).
fof(t113_zfmisc_1, conjecture, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/579d6d627e734c97be966e07f621e954.E_input', t113_zfmisc_1)).
fof(c_0_4, plain, ![X30, X31, X32, X33, X34, X35, X36, X37]:((((r2_hidden(X33,X30)|~r2_hidden(X33,X32)|X32!=k4_xboole_0(X30,X31))&(~r2_hidden(X33,X31)|~r2_hidden(X33,X32)|X32!=k4_xboole_0(X30,X31)))&(~r2_hidden(X34,X30)|r2_hidden(X34,X31)|r2_hidden(X34,X32)|X32!=k4_xboole_0(X30,X31)))&((~r2_hidden(esk9_3(X35,X36,X37),X37)|(~r2_hidden(esk9_3(X35,X36,X37),X35)|r2_hidden(esk9_3(X35,X36,X37),X36))|X37=k4_xboole_0(X35,X36))&((r2_hidden(esk9_3(X35,X36,X37),X35)|r2_hidden(esk9_3(X35,X36,X37),X37)|X37=k4_xboole_0(X35,X36))&(~r2_hidden(esk9_3(X35,X36,X37),X36)|r2_hidden(esk9_3(X35,X36,X37),X37)|X37=k4_xboole_0(X35,X36))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_5, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_6, plain, ![X29]:k4_xboole_0(k1_xboole_0,X29)=k1_xboole_0, inference(variable_rename,[status(thm)],[t4_boole])).
cnf(c_0_7, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_5])).
cnf(c_0_8, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, plain, (r2_hidden(esk9_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk9_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10, plain, (r2_hidden(esk9_3(X1,X2,X3),X1)|r2_hidden(esk9_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_11, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_12, plain, (X1=k4_xboole_0(X2,X2)|r2_hidden(esk9_3(X2,X2,X1),X1)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_13, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_14, plain, (k4_xboole_0(X1,X1)=k1_xboole_0|~r2_hidden(esk9_3(X1,X1,k1_xboole_0),X2)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
fof(c_0_15, plain, ![X9, X10, X11, X12, X15, X16, X17, X18, X19, X20, X22, X23]:((((r2_hidden(esk3_4(X9,X10,X11,X12),X9)|~r2_hidden(X12,X11)|X11!=k2_zfmisc_1(X9,X10))&((r2_hidden(esk4_4(X9,X10,X11,X12),X10)|~r2_hidden(X12,X11)|X11!=k2_zfmisc_1(X9,X10))&(X12=k4_tarski(esk3_4(X9,X10,X11,X12),esk4_4(X9,X10,X11,X12))|~r2_hidden(X12,X11)|X11!=k2_zfmisc_1(X9,X10))))&(~r2_hidden(X16,X9)|(~r2_hidden(X17,X10)|X15!=k4_tarski(X16,X17))|r2_hidden(X15,X11)|X11!=k2_zfmisc_1(X9,X10)))&((~r2_hidden(esk5_3(X18,X19,X20),X20)|(~r2_hidden(X22,X18)|(~r2_hidden(X23,X19)|esk5_3(X18,X19,X20)!=k4_tarski(X22,X23)))|X20=k2_zfmisc_1(X18,X19))&((r2_hidden(esk6_3(X18,X19,X20),X18)|r2_hidden(esk5_3(X18,X19,X20),X20)|X20=k2_zfmisc_1(X18,X19))&((r2_hidden(esk7_3(X18,X19,X20),X19)|r2_hidden(esk5_3(X18,X19,X20),X20)|X20=k2_zfmisc_1(X18,X19))&(esk5_3(X18,X19,X20)=k4_tarski(esk6_3(X18,X19,X20),esk7_3(X18,X19,X20))|r2_hidden(esk5_3(X18,X19,X20),X20)|X20=k2_zfmisc_1(X18,X19)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
cnf(c_0_16, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_13])).
cnf(c_0_17, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_14, c_0_12])).
cnf(c_0_18, plain, (r2_hidden(esk3_4(X1,X2,X3,X4),X1)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, plain, (~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_11])).
cnf(c_0_20, plain, (r2_hidden(esk3_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X1)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_18])).
cnf(c_0_21, plain, (~r2_hidden(X1,k2_zfmisc_1(k1_xboole_0,X2))), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_22, plain, (X1=k1_xboole_0|r2_hidden(esk9_3(X2,X2,X1),X1)), inference(rw,[status(thm)],[c_0_12, c_0_17])).
cnf(c_0_23, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_24, plain, (r2_hidden(esk4_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_25, negated_conjecture, ~(![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0))), inference(assume_negation,[status(cth)],[t113_zfmisc_1])).
cnf(c_0_26, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_23]), c_0_19])).
cnf(c_0_27, plain, (r2_hidden(esk4_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X2)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_24])).
fof(c_0_28, negated_conjecture, (((esk1_0!=k1_xboole_0|k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0)&(esk2_0!=k1_xboole_0|k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0))&(k2_zfmisc_1(esk1_0,esk2_0)=k1_xboole_0|(esk1_0=k1_xboole_0|esk2_0=k1_xboole_0))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])])).
cnf(c_0_29, plain, (X1!=k1_xboole_0|~r2_hidden(X2,k2_zfmisc_1(X3,X1))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_30, plain, (X1=k1_xboole_0|r2_hidden(esk9_3(k1_xboole_0,X2,X1),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_10]), c_0_8])).
cnf(c_0_31, negated_conjecture, (esk2_0!=k1_xboole_0|k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_32, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X2!=k1_xboole_0), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_33, plain, (X1!=k1_xboole_0|~r2_hidden(X2,k2_zfmisc_1(X1,X3))), inference(spm,[status(thm)],[c_0_26, c_0_20])).
cnf(c_0_34, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k1_xboole_0|esk1_0=k1_xboole_0|esk2_0=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, negated_conjecture, (esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_36, negated_conjecture, (esk1_0!=k1_xboole_0|k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_33, c_0_30])).
cnf(c_0_38, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k1_xboole_0|esk1_0=k1_xboole_0), inference(sr,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_39, negated_conjecture, (esk1_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_36, c_0_37])).
fof(c_0_40, plain, (~epred4_0<=>![X2]:X2!=k1_xboole_0), introduced(definition)).
fof(c_0_41, plain, (~epred3_0<=>![X4, X3, X1]:(~r2_hidden(X3,esk1_0)|~r2_hidden(X4,esk2_0)|X1!=k4_tarski(X3,X4))), introduced(definition)).
cnf(c_0_42, plain, (r2_hidden(X5,X6)|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)|X5!=k4_tarski(X1,X3)|X6!=k2_zfmisc_1(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_43, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k1_xboole_0), inference(sr,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_44, negated_conjecture, (epred4_0|X1!=k1_xboole_0), inference(split_equiv,[status(thm)],[c_0_40])).
cnf(c_0_45, negated_conjecture, (~epred4_0|~epred3_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_26]), c_0_41]), c_0_40])).
cnf(c_0_46, negated_conjecture, (epred4_0), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_47, negated_conjecture, (~epred3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_46])])).
fof(c_0_48, plain, (~epred5_0<=>![X1]:~r2_hidden(X1,esk1_0)), introduced(definition)).
fof(c_0_49, plain, (~epred6_0<=>![X2]:~r2_hidden(X2,esk2_0)), introduced(definition)).
cnf(c_0_50, negated_conjecture, (X1!=k4_tarski(X2,X3)|~r2_hidden(X2,esk1_0)|~r2_hidden(X3,esk2_0)), inference(sr,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_41]), c_0_47])).
cnf(c_0_51, negated_conjecture, (epred5_0|~r2_hidden(X1,esk1_0)), inference(split_equiv,[status(thm)],[c_0_48])).
cnf(c_0_52, negated_conjecture, (~epred6_0|~epred5_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(er,[status(thm)],[c_0_50]), c_0_48]), c_0_49])).
cnf(c_0_53, negated_conjecture, (epred5_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_30]), c_0_39])).
cnf(c_0_54, negated_conjecture, (~epred6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_53])])).
cnf(c_0_55, negated_conjecture, (~r2_hidden(X1,esk2_0)), inference(sr,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_49]), c_0_54])).
cnf(c_0_56, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_30]), c_0_35]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 57
# Proof object clause steps            : 44
# Proof object formula steps           : 13
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 24
# Proof object simplifying inferences  : 19
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 20
# Processed clauses                    : 6340
# ...of these trivial                  : 14
# ...subsumed                          : 5546
# ...remaining for further processing  : 780
# Other redundant clauses eliminated   : 2100
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 48
# Backward-rewritten                   : 12
# Generated clauses                    : 196945
# ...of the previous two non-trivial   : 188696
# Contextual simplify-reflections      : 43
# Paramodulations                      : 194241
# Factorizations                       : 351
# Equation resolutions                 : 2341
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 694
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 677
# Current number of unprocessed clauses: 182283
# ...number of literals in the above   : 1743602
# Current number of archived formulas  : 0
# Current number of archived clauses   : 83
# Clause-clause subsumption calls (NU) : 548863
# Rec. Clause-clause subsumption calls : 106295
# Non-unit clause-clause subsumptions  : 3616
# Unit Clause-clause subsumption calls : 1694
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 16
# BW rewrite match successes           : 8
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5765157

# -------------------------------------------------
# User time                : 5.778 s
# System time              : 0.053 s
# Total time               : 5.831 s
# Maximum resident set size: 174332 pages
