# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/57b27e340fb943c3a9597181aeeab545.E_input', t3_xboole_0)).
fof(t166_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k10_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k2_relat_1(X3))&(r2_hidden(k4_tarski(X1,X4),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/57b27e340fb943c3a9597181aeeab545.E_input', t166_relat_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/57b27e340fb943c3a9597181aeeab545.E_input', t7_xboole_0)).
fof(t173_relat_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>(k10_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k2_relat_1(X2),X1))), file('data/MPTP/proofs/57b27e340fb943c3a9597181aeeab545.E_input', t173_relat_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/57b27e340fb943c3a9597181aeeab545.E_input', d5_relat_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/57b27e340fb943c3a9597181aeeab545.E_input', d7_xboole_0)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/57b27e340fb943c3a9597181aeeab545.E_input', t65_xboole_1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/57b27e340fb943c3a9597181aeeab545.E_input', t4_xboole_0)).
fof(c_0_8, plain, ![X5, X6, X8, X9, X10]:(((r2_hidden(esk1_2(X5,X6),X5)|r1_xboole_0(X5,X6))&(r2_hidden(esk1_2(X5,X6),X6)|r1_xboole_0(X5,X6)))&(~r2_hidden(X10,X8)|~r2_hidden(X10,X9)|~r1_xboole_0(X8,X9))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_9, plain, ![X11, X12, X13, X15]:(((r2_hidden(esk2_3(X11,X12,X13),k2_relat_1(X13))|~r2_hidden(X11,k10_relat_1(X13,X12))|~v1_relat_1(X13))&((r2_hidden(k4_tarski(X11,esk2_3(X11,X12,X13)),X13)|~r2_hidden(X11,k10_relat_1(X13,X12))|~v1_relat_1(X13))&(r2_hidden(esk2_3(X11,X12,X13),X12)|~r2_hidden(X11,k10_relat_1(X13,X12))|~v1_relat_1(X13))))&(~r2_hidden(X15,k2_relat_1(X13))|(~r2_hidden(k4_tarski(X11,X15),X13)|~r2_hidden(X15,X12))|r2_hidden(X11,k10_relat_1(X13,X12))|~v1_relat_1(X13))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t166_relat_1])])])])])).
fof(c_0_10, plain, ![X16]:(X16=k1_xboole_0|r2_hidden(esk3_1(X16),X16)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_11, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (r2_hidden(esk1_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_13, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>(k10_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k2_relat_1(X2),X1)))), inference(assume_negation,[status(cth)],[t173_relat_1])).
cnf(c_0_14, plain, (r2_hidden(esk2_3(X1,X2,X3),k2_relat_1(X3))|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (r2_hidden(esk2_3(X1,X2,X3),X2)|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, plain, (X1=k1_xboole_0|r2_hidden(esk3_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, plain, (r1_xboole_0(X1,X2)|~r2_hidden(esk1_2(X1,X2),X3)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_18, plain, (r2_hidden(esk1_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_19, negated_conjecture, (v1_relat_1(esk5_0)&((k10_relat_1(esk5_0,esk4_0)!=k1_xboole_0|~r1_xboole_0(k2_relat_1(esk5_0),esk4_0))&(k10_relat_1(esk5_0,esk4_0)=k1_xboole_0|r1_xboole_0(k2_relat_1(esk5_0),esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
fof(c_0_20, plain, ![X20, X21, X22, X24, X25, X26, X27, X29]:(((~r2_hidden(X22,X21)|r2_hidden(k4_tarski(esk6_3(X20,X21,X22),X22),X20)|X21!=k2_relat_1(X20))&(~r2_hidden(k4_tarski(X25,X24),X20)|r2_hidden(X24,X21)|X21!=k2_relat_1(X20)))&((~r2_hidden(esk7_2(X26,X27),X27)|~r2_hidden(k4_tarski(X29,esk7_2(X26,X27)),X26)|X27=k2_relat_1(X26))&(r2_hidden(esk7_2(X26,X27),X27)|r2_hidden(k4_tarski(esk8_2(X26,X27),esk7_2(X26,X27)),X26)|X27=k2_relat_1(X26)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
fof(c_0_21, plain, ![X31, X32]:((~r1_xboole_0(X31,X32)|k3_xboole_0(X31,X32)=k1_xboole_0)&(k3_xboole_0(X31,X32)!=k1_xboole_0|r1_xboole_0(X31,X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
fof(c_0_22, plain, ![X33]:r1_xboole_0(X33,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_23, plain, (~v1_relat_1(X1)|~r2_hidden(esk2_3(X2,X3,X1),X4)|~r2_hidden(X2,k10_relat_1(X1,X3))|~r1_xboole_0(X4,k2_relat_1(X1))), inference(spm,[status(thm)],[c_0_11, c_0_14])).
cnf(c_0_24, plain, (k10_relat_1(X1,X2)=k1_xboole_0|r2_hidden(esk2_3(esk3_1(k10_relat_1(X1,X2)),X2,X1),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_25, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_26, negated_conjecture, (k10_relat_1(esk5_0,esk4_0)=k1_xboole_0|r1_xboole_0(k2_relat_1(esk5_0),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (r2_hidden(k4_tarski(esk6_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_28, plain, ![X34, X35, X37, X38, X39]:((r1_xboole_0(X34,X35)|r2_hidden(esk9_2(X34,X35),k3_xboole_0(X34,X35)))&(~r2_hidden(X39,k3_xboole_0(X37,X38))|~r1_xboole_0(X37,X38))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
cnf(c_0_29, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, negated_conjecture, (k10_relat_1(esk5_0,esk4_0)!=k1_xboole_0|~r1_xboole_0(k2_relat_1(esk5_0),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_32, plain, (k10_relat_1(X1,X2)=k1_xboole_0|~v1_relat_1(X1)|~r1_xboole_0(X2,k2_relat_1(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_16])).
cnf(c_0_33, negated_conjecture, (k10_relat_1(esk5_0,esk4_0)=k1_xboole_0|r1_xboole_0(esk4_0,k2_relat_1(esk5_0))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_34, negated_conjecture, (v1_relat_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_35, plain, (r2_hidden(X3,k10_relat_1(X2,X4))|~r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)|~r2_hidden(X1,X4)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_36, plain, (r2_hidden(k4_tarski(esk6_3(X1,k2_relat_1(X1),X2),X2),X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_39, negated_conjecture, (r2_hidden(esk1_2(k2_relat_1(esk5_0),esk4_0),esk4_0)|k10_relat_1(esk5_0,esk4_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_31, c_0_18])).
cnf(c_0_40, negated_conjecture, (k10_relat_1(esk5_0,esk4_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])])).
cnf(c_0_41, plain, (r2_hidden(esk6_3(X1,k2_relat_1(X1),X2),k10_relat_1(X1,X3))|~v1_relat_1(X1)|~r2_hidden(X2,k2_relat_1(X1))|~r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_42, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_30])])).
cnf(c_0_43, negated_conjecture, (r2_hidden(esk1_2(k2_relat_1(esk5_0),esk4_0),esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_40])])).
cnf(c_0_44, negated_conjecture, (~r1_xboole_0(k2_relat_1(esk5_0),esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_40])])).
cnf(c_0_45, negated_conjecture, (~r2_hidden(X1,k2_relat_1(esk5_0))|~r2_hidden(X1,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_40]), c_0_34])]), c_0_42])).
cnf(c_0_46, plain, (r2_hidden(esk1_2(X1,X2),X1)|r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_25, c_0_12])).
cnf(c_0_47, negated_conjecture, (~r1_xboole_0(esk4_0,k2_relat_1(esk5_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_43]), c_0_44])).
cnf(c_0_48, negated_conjecture, (r1_xboole_0(X1,k2_relat_1(esk5_0))|~r2_hidden(esk1_2(X1,k2_relat_1(esk5_0)),esk4_0)), inference(spm,[status(thm)],[c_0_45, c_0_18])).
cnf(c_0_49, negated_conjecture, (r2_hidden(esk1_2(esk4_0,k2_relat_1(esk5_0)),esk4_0)), inference(spm,[status(thm)],[c_0_44, c_0_46])).
cnf(c_0_50, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 51
# Proof object clause steps            : 34
# Proof object formula steps           : 17
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 15
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 20
# Processed clauses                    : 5219
# ...of these trivial                  : 26
# ...subsumed                          : 4318
# ...remaining for further processing  : 875
# Other redundant clauses eliminated   : 125
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 31
# Generated clauses                    : 81321
# ...of the previous two non-trivial   : 58646
# Contextual simplify-reflections      : 56
# Paramodulations                      : 81076
# Factorizations                       : 68
# Equation resolutions                 : 177
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 818
#    Positive orientable unit clauses  : 18
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 18
#    Non-unit-clauses                  : 782
# Current number of unprocessed clauses: 53366
# ...number of literals in the above   : 334054
# Current number of archived formulas  : 0
# Current number of archived clauses   : 57
# Clause-clause subsumption calls (NU) : 213865
# Rec. Clause-clause subsumption calls : 145368
# Non-unit clause-clause subsumptions  : 3669
# Unit Clause-clause subsumption calls : 69
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1862006

# -------------------------------------------------
# User time                : 1.162 s
# System time              : 0.007 s
# Total time               : 1.169 s
# Maximum resident set size: 50808 pages
