# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', t70_enumset1)).
fof(t95_xboole_1, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2)), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', t95_xboole_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', t72_enumset1)).
fof(t21_xboole_1, axiom, ![X1, X2]:k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', t21_xboole_1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', t73_enumset1)).
fof(t47_zfmisc_1, conjecture, ![X1, X2, X3]:(r1_tarski(k2_xboole_0(k2_tarski(X1,X2),X3),X3)=>r2_hidden(X1,X3)), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', t47_zfmisc_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', idempotence_k3_xboole_0)).
fof(idempotence_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', idempotence_k2_xboole_0)).
fof(t46_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0, file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', t46_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', t100_xboole_1)).
fof(t91_xboole_1, axiom, ![X1, X2, X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3)), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', t91_xboole_1)).
fof(t102_enumset1, axiom, ![X1, X2, X3]:k1_enumset1(X1,X2,X3)=k1_enumset1(X3,X2,X1), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', t102_enumset1)).
fof(t43_enumset1, axiom, ![X1, X2, X3]:k1_enumset1(X1,X2,X3)=k2_xboole_0(k2_tarski(X1,X2),k1_tarski(X3)), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', t43_enumset1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', t69_enumset1)).
fof(l53_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X4)), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', l53_enumset1)).
fof(t60_xboole_1, axiom, ![X1, X2]:~((r1_tarski(X1,X2)&r2_xboole_0(X2,X1))), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', t60_xboole_1)).
fof(d8_xboole_0, axiom, ![X1, X2]:(r2_xboole_0(X1,X2)<=>(r1_tarski(X1,X2)&~(X1=X2))), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', d8_xboole_0)).
fof(commutativity_k5_xboole_0, axiom, ![X1, X2]:k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', commutativity_k5_xboole_0)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', t17_xboole_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', t7_xboole_1)).
fof(d1_enumset1, axiom, ![X1, X2, X3, X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3)))))), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', d1_enumset1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/582955d6645a43a1892a66de6efc8bcf.E_input', d3_tarski)).
fof(c_0_24, plain, ![X19, X20]:k3_tarski(k2_tarski(X19,X20))=k2_xboole_0(X19,X20), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_25, plain, ![X40, X41]:k1_enumset1(X40,X40,X41)=k2_tarski(X40,X41), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_26, plain, ![X12, X13]:k3_xboole_0(X12,X13)=k5_xboole_0(k5_xboole_0(X12,X13),k2_xboole_0(X12,X13)), inference(variable_rename,[status(thm)],[t95_xboole_1])).
cnf(c_0_27, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_28, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_29, plain, ![X14, X15, X16]:k2_enumset1(X14,X14,X15,X16)=k1_enumset1(X14,X15,X16), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_30, plain, ![X60, X61, X62, X63]:k3_enumset1(X60,X60,X61,X62,X63)=k2_enumset1(X60,X61,X62,X63), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_31, plain, ![X69, X70]:k3_xboole_0(X69,k2_xboole_0(X69,X70))=X69, inference(variable_rename,[status(thm)],[t21_xboole_1])).
cnf(c_0_32, plain, (k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_35, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_36, plain, ![X23, X24, X25, X26, X27]:k4_enumset1(X23,X23,X24,X25,X26,X27)=k3_enumset1(X23,X24,X25,X26,X27), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_37, negated_conjecture, ~(![X1, X2, X3]:(r1_tarski(k2_xboole_0(k2_tarski(X1,X2),X3),X3)=>r2_hidden(X1,X3))), inference(assume_negation,[status(cth)],[t47_zfmisc_1])).
fof(c_0_38, plain, ![X9]:k3_xboole_0(X9,X9)=X9, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
fof(c_0_39, plain, ![X72]:k2_xboole_0(X72,X72)=X72, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k2_xboole_0])])).
fof(c_0_40, plain, ![X21, X22]:k4_xboole_0(X21,k2_xboole_0(X21,X22))=k1_xboole_0, inference(variable_rename,[status(thm)],[t46_xboole_1])).
fof(c_0_41, plain, ![X64, X65]:k4_xboole_0(X64,X65)=k5_xboole_0(X64,k3_xboole_0(X64,X65)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_42, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_43, plain, (k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k3_enumset1(X1,X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_33]), c_0_34]), c_0_35])).
cnf(c_0_44, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_45, plain, ![X73, X74, X75]:k5_xboole_0(k5_xboole_0(X73,X74),X75)=k5_xboole_0(X73,k5_xboole_0(X74,X75)), inference(variable_rename,[status(thm)],[t91_xboole_1])).
fof(c_0_46, negated_conjecture, (r1_tarski(k2_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0),esk3_0)&~r2_hidden(esk1_0,esk3_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_37])])])).
fof(c_0_47, plain, ![X46, X47, X48]:k1_enumset1(X46,X47,X48)=k1_enumset1(X48,X47,X46), inference(variable_rename,[status(thm)],[t102_enumset1])).
fof(c_0_48, plain, ![X66, X67, X68]:k1_enumset1(X66,X67,X68)=k2_xboole_0(k2_tarski(X66,X67),k1_tarski(X68)), inference(variable_rename,[status(thm)],[t43_enumset1])).
fof(c_0_49, plain, ![X71]:k2_tarski(X71,X71)=k1_tarski(X71), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_50, plain, ![X36, X37, X38, X39]:k2_enumset1(X36,X37,X38,X39)=k2_xboole_0(k2_tarski(X36,X37),k2_tarski(X38,X39)), inference(variable_rename,[status(thm)],[l53_enumset1])).
cnf(c_0_51, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_52, plain, (k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_53, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_54, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_55, plain, (k5_xboole_0(k5_xboole_0(X1,k3_tarski(k4_enumset1(X1,X1,X1,X1,X1,X2))),k3_tarski(k4_enumset1(X1,X1,X1,X1,X1,k3_tarski(k4_enumset1(X1,X1,X1,X1,X1,X2)))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_33]), c_0_34]), c_0_35]), c_0_43]), c_0_44]), c_0_44]), c_0_44])).
cnf(c_0_56, plain, (k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_57, negated_conjecture, (r1_tarski(k2_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_58, plain, (k1_enumset1(X1,X2,X3)=k1_enumset1(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_59, plain, (k1_enumset1(X1,X2,X3)=k2_xboole_0(k2_tarski(X1,X2),k1_tarski(X3))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_60, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_61, plain, (k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X4))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_62, plain, (k5_xboole_0(k5_xboole_0(X1,X1),k3_tarski(k4_enumset1(X1,X1,X1,X1,X1,X1)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_43]), c_0_44])).
cnf(c_0_63, plain, (k3_tarski(k4_enumset1(X1,X1,X1,X1,X1,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_33]), c_0_34]), c_0_35]), c_0_44])).
cnf(c_0_64, plain, (k5_xboole_0(X1,k5_xboole_0(k5_xboole_0(X1,k3_tarski(k4_enumset1(X1,X1,X1,X1,X1,X2))),k3_tarski(k4_enumset1(X1,X1,X1,X1,X1,k3_tarski(k4_enumset1(X1,X1,X1,X1,X1,X2))))))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_33]), c_0_54]), c_0_34]), c_0_35]), c_0_43]), c_0_44]), c_0_44]), c_0_44])).
cnf(c_0_65, plain, (k5_xboole_0(X1,k5_xboole_0(k3_tarski(k4_enumset1(X1,X1,X1,X1,X1,X2)),k3_tarski(k4_enumset1(X1,X1,X1,X1,X1,k3_tarski(k4_enumset1(X1,X1,X1,X1,X1,X2))))))=X1), inference(rw,[status(thm)],[c_0_55, c_0_56])).
fof(c_0_66, plain, ![X10, X11]:(~r1_tarski(X10,X11)|~r2_xboole_0(X11,X10)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t60_xboole_1])])).
fof(c_0_67, plain, ![X44, X45]:(((r1_tarski(X44,X45)|~r2_xboole_0(X44,X45))&(X44!=X45|~r2_xboole_0(X44,X45)))&(~r1_tarski(X44,X45)|X44=X45|r2_xboole_0(X44,X45))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_xboole_0])])])])).
cnf(c_0_68, negated_conjecture, (r1_tarski(k3_tarski(k4_enumset1(k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0),k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0),k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0),k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0),k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0),esk3_0)),esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_28]), c_0_33]), c_0_34]), c_0_34]), c_0_34]), c_0_35]), c_0_35]), c_0_35]), c_0_35]), c_0_44]), c_0_44]), c_0_44]), c_0_44]), c_0_44])).
cnf(c_0_69, plain, (k4_enumset1(X1,X1,X1,X1,X2,X3)=k4_enumset1(X3,X3,X3,X3,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_34]), c_0_34]), c_0_35]), c_0_35]), c_0_44]), c_0_44])).
cnf(c_0_70, plain, (k4_enumset1(X1,X1,X1,X1,X2,X3)=k3_tarski(k4_enumset1(k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X3,X3,X3,X3,X3,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_60]), c_0_28]), c_0_28]), c_0_33]), c_0_34]), c_0_34]), c_0_34]), c_0_34]), c_0_34]), c_0_35]), c_0_35]), c_0_35]), c_0_35]), c_0_35]), c_0_35]), c_0_44]), c_0_44]), c_0_44]), c_0_44]), c_0_44]), c_0_44]), c_0_44])).
cnf(c_0_71, plain, (k4_enumset1(X1,X1,X1,X2,X3,X4)=k3_tarski(k4_enumset1(k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X1,X1,X1,X1,X1,X2),k4_enumset1(X3,X3,X3,X3,X3,X4)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_28]), c_0_28]), c_0_33]), c_0_34]), c_0_34]), c_0_34]), c_0_34]), c_0_35]), c_0_35]), c_0_35]), c_0_35]), c_0_35]), c_0_35]), c_0_44]), c_0_44]), c_0_44]), c_0_44]), c_0_44]), c_0_44]), c_0_44])).
fof(c_0_72, plain, ![X28, X29]:k5_xboole_0(X28,X29)=k5_xboole_0(X29,X28), inference(variable_rename,[status(thm)],[commutativity_k5_xboole_0])).
cnf(c_0_73, plain, (k5_xboole_0(X1,k5_xboole_0(X1,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_63]), c_0_56])).
cnf(c_0_74, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_56]), c_0_65])).
fof(c_0_75, plain, ![X17, X18]:r1_tarski(k3_xboole_0(X17,X18),X17), inference(variable_rename,[status(thm)],[t17_xboole_1])).
cnf(c_0_76, plain, (~r1_tarski(X1,X2)|~r2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_77, plain, (X1=X2|r2_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_78, negated_conjecture, (r1_tarski(k3_tarski(k4_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0),k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0))),esk3_0)), inference(rw,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_79, plain, (k4_enumset1(X1,X1,X1,X2,X3,X3)=k4_enumset1(X1,X1,X1,X1,X2,X3)), inference(rw,[status(thm)],[c_0_70, c_0_71])).
fof(c_0_80, plain, ![X42, X43]:r1_tarski(X42,k2_xboole_0(X42,X43)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_81, plain, (k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_82, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_83, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_84, plain, (X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_85, negated_conjecture, (r1_tarski(k3_tarski(k4_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0))),esk3_0)), inference(rw,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_86, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_87, plain, (k5_xboole_0(k1_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_88, plain, (r1_tarski(k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k4_enumset1(X1,X1,X1,X1,X1,X2))),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_83, c_0_43]), c_0_44])).
cnf(c_0_89, negated_conjecture, (k3_tarski(k4_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)))=esk3_0|~r1_tarski(esk3_0,k3_tarski(k4_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0))))), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_90, plain, (r1_tarski(X1,k3_tarski(k4_enumset1(X1,X1,X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_33]), c_0_34]), c_0_35]), c_0_44])).
cnf(c_0_91, plain, (k5_xboole_0(X1,k5_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_74]), c_0_87])).
fof(c_0_92, plain, ![X49, X50, X51, X52, X53, X54, X55, X56, X57, X58]:(((~r2_hidden(X53,X52)|(X53=X49|(X53=X50|X53=X51))|X52!=k1_enumset1(X49,X50,X51))&((X54!=X49|r2_hidden(X54,X52)|X52!=k1_enumset1(X49,X50,X51))&((X54!=X50|r2_hidden(X54,X52)|X52!=k1_enumset1(X49,X50,X51))&(X54!=X51|r2_hidden(X54,X52)|X52!=k1_enumset1(X49,X50,X51)))))&(((esk5_4(X55,X56,X57,X58)!=X55|~r2_hidden(esk5_4(X55,X56,X57,X58),X58)|X58=k1_enumset1(X55,X56,X57))&((esk5_4(X55,X56,X57,X58)!=X56|~r2_hidden(esk5_4(X55,X56,X57,X58),X58)|X58=k1_enumset1(X55,X56,X57))&(esk5_4(X55,X56,X57,X58)!=X57|~r2_hidden(esk5_4(X55,X56,X57,X58),X58)|X58=k1_enumset1(X55,X56,X57))))&(r2_hidden(esk5_4(X55,X56,X57,X58),X58)|(esk5_4(X55,X56,X57,X58)=X55|(esk5_4(X55,X56,X57,X58)=X56|esk5_4(X55,X56,X57,X58)=X57))|X58=k1_enumset1(X55,X56,X57)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])).
fof(c_0_93, plain, ![X30, X31, X32, X33, X34]:((~r1_tarski(X30,X31)|(~r2_hidden(X32,X30)|r2_hidden(X32,X31)))&((r2_hidden(esk4_2(X33,X34),X33)|r1_tarski(X33,X34))&(~r2_hidden(esk4_2(X33,X34),X34)|r1_tarski(X33,X34)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_94, plain, (r1_tarski(k5_xboole_0(X1,k5_xboole_0(X2,k3_tarski(k4_enumset1(X1,X1,X1,X1,X1,X2)))),X1)), inference(rw,[status(thm)],[c_0_88, c_0_56])).
cnf(c_0_95, negated_conjecture, (k3_tarski(k4_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)))=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_90])])).
cnf(c_0_96, plain, (k5_xboole_0(X1,k5_xboole_0(X2,X1))=X2), inference(spm,[status(thm)],[c_0_91, c_0_81])).
cnf(c_0_97, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X2,X5)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_98, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_99, negated_conjecture, (r1_tarski(k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0),esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_96])).
cnf(c_0_100, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k4_enumset1(X4,X4,X4,X4,X2,X5)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_34]), c_0_35]), c_0_44])).
cnf(c_0_101, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_102, plain, (r2_hidden(X1,k4_enumset1(X2,X2,X2,X2,X1,X3))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_100])])).
cnf(c_0_103, negated_conjecture, (~r2_hidden(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_104, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_103]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 105
# Proof object clause steps            : 56
# Proof object formula steps           : 49
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 24
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 113
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 24
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 36
# Removed in clause preprocessing      : 8
# Initial clauses in saturation        : 28
# Processed clauses                    : 74
# ...of these trivial                  : 2
# ...subsumed                          : 2
# ...remaining for further processing  : 70
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 7
# Generated clauses                    : 71
# ...of the previous two non-trivial   : 46
# Contextual simplify-reflections      : 0
# Paramodulations                      : 66
# Factorizations                       : 0
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 30
#    Positive orientable unit clauses  : 15
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 12
# Current number of unprocessed clauses: 25
# ...number of literals in the above   : 38
# Current number of archived formulas  : 0
# Current number of archived clauses   : 43
# Clause-clause subsumption calls (NU) : 18
# Rec. Clause-clause subsumption calls : 18
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 7
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 75
# BW rewrite match successes           : 36
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2557

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.004 s
# Total time               : 0.011 s
# Maximum resident set size: 28808 pages
