# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', t3_subset)).
fof(t63_subset_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', t63_subset_1)).
fof(t44_tex_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v2_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v2_tex_2(X2,X1)<=>v1_zfmisc_1(X2)))), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', t44_tex_2)).
fof(t1_tex_2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((~(v1_xboole_0(X2))&v1_zfmisc_1(X2))=>(r1_tarski(X1,X2)=>X1=X2))), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', t1_tex_2)).
fof(fc2_xboole_0, axiom, ![X1]:~(v1_xboole_0(k1_tarski(X1))), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', fc2_xboole_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', d1_xboole_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', t4_subset)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', cc1_subset_1)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', redefinition_k6_domain_1)).
fof(t34_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>~((v2_tex_2(X2,X1)&![X3]:((~(v2_struct_0(X3))&(v1_pre_topc(X3)&(v1_tdlat_3(X3)&m1_pre_topc(X3,X1))))=>~(X2=u1_struct_0(X3))))))), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', t34_tex_2)).
fof(t36_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>v2_tex_2(k6_domain_1(u1_struct_0(X1),X2),X1))), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', t36_tex_2)).
fof(rc1_subset_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>?[X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))&~(v1_xboole_0(X2)))), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', rc1_subset_1)).
fof(cc6_tdlat_3, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v2_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:(m1_pre_topc(X2,X1)=>v2_tdlat_3(X2))), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', cc6_tdlat_3)).
fof(cc2_tdlat_3, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_tdlat_3(X1)=>v2_pre_topc(X1))), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', cc2_tdlat_3)).
fof(cc2_tex_1, axiom, ![X1]:(l1_pre_topc(X1)=>((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v1_tdlat_3(X1)&v2_tdlat_3(X1))))=>(~(v2_struct_0(X1))&(v7_struct_0(X1)&v2_pre_topc(X1))))), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', cc2_tex_1)).
fof(cc1_tdlat_3, axiom, ![X1]:(l1_pre_topc(X1)=>(v1_tdlat_3(X1)=>v2_pre_topc(X1))), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', cc1_tdlat_3)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', dt_m1_pre_topc)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', t1_xboole_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', d3_tarski)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', dt_l1_pre_topc)).
fof(fc7_struct_0, axiom, ![X1]:((v7_struct_0(X1)&l1_struct_0(X1))=>v1_zfmisc_1(u1_struct_0(X1))), file('data/MPTP/proofs/59fd5912b8fc4918a7ded686b602066e.E_input', fc7_struct_0)).
fof(c_0_21, plain, ![X14, X15]:((~m1_subset_1(X14,k1_zfmisc_1(X15))|r1_tarski(X14,X15))&(~r1_tarski(X14,X15)|m1_subset_1(X14,k1_zfmisc_1(X15)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_22, plain, ![X16, X17]:(~r2_hidden(X16,X17)|m1_subset_1(k1_tarski(X16),k1_zfmisc_1(X17))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_subset_1])])).
fof(c_0_23, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v2_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v2_tex_2(X2,X1)<=>v1_zfmisc_1(X2))))), inference(assume_negation,[status(cth)],[t44_tex_2])).
fof(c_0_24, plain, ![X36, X37]:(v1_xboole_0(X36)|(v1_xboole_0(X37)|~v1_zfmisc_1(X37)|(~r1_tarski(X36,X37)|X36=X37))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t1_tex_2])])])])).
cnf(c_0_25, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_26, plain, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_27, plain, ![X22]:~v1_xboole_0(k1_tarski(X22)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_xboole_0])])).
fof(c_0_28, plain, ![X41, X42, X43]:((~v1_xboole_0(X41)|~r2_hidden(X42,X41))&(r2_hidden(esk6_1(X43),X43)|v1_xboole_0(X43))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_29, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&(v2_tdlat_3(esk1_0)&l1_pre_topc(esk1_0))))&((~v1_xboole_0(esk2_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))&((~v2_tex_2(esk2_0,esk1_0)|~v1_zfmisc_1(esk2_0))&(v2_tex_2(esk2_0,esk1_0)|v1_zfmisc_1(esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_23])])])])).
fof(c_0_30, plain, ![X25, X26, X27]:(~r2_hidden(X25,X26)|~m1_subset_1(X26,k1_zfmisc_1(X27))|m1_subset_1(X25,X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_31, plain, ![X32, X33]:(~v1_xboole_0(X32)|(~m1_subset_1(X33,k1_zfmisc_1(X32))|v1_xboole_0(X33))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
cnf(c_0_32, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|X1=X2|~v1_zfmisc_1(X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_34, plain, (~v1_xboole_0(k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, plain, (r2_hidden(esk6_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_38, plain, ![X20, X21]:(v1_xboole_0(X20)|~m1_subset_1(X21,X20)|k6_domain_1(X20,X21)=k1_tarski(X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_39, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_41, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_42, plain, ![X28, X29]:(((~v2_struct_0(esk4_2(X28,X29))|~v2_tex_2(X29,X28)|(v1_xboole_0(X29)|~m1_subset_1(X29,k1_zfmisc_1(u1_struct_0(X28))))|(v2_struct_0(X28)|(~v2_pre_topc(X28)|~l1_pre_topc(X28))))&((v1_pre_topc(esk4_2(X28,X29))|~v2_tex_2(X29,X28)|(v1_xboole_0(X29)|~m1_subset_1(X29,k1_zfmisc_1(u1_struct_0(X28))))|(v2_struct_0(X28)|(~v2_pre_topc(X28)|~l1_pre_topc(X28))))&((v1_tdlat_3(esk4_2(X28,X29))|~v2_tex_2(X29,X28)|(v1_xboole_0(X29)|~m1_subset_1(X29,k1_zfmisc_1(u1_struct_0(X28))))|(v2_struct_0(X28)|(~v2_pre_topc(X28)|~l1_pre_topc(X28))))&(m1_pre_topc(esk4_2(X28,X29),X28)|~v2_tex_2(X29,X28)|(v1_xboole_0(X29)|~m1_subset_1(X29,k1_zfmisc_1(u1_struct_0(X28))))|(v2_struct_0(X28)|(~v2_pre_topc(X28)|~l1_pre_topc(X28)))))))&(X29=u1_struct_0(esk4_2(X28,X29))|~v2_tex_2(X29,X28)|(v1_xboole_0(X29)|~m1_subset_1(X29,k1_zfmisc_1(u1_struct_0(X28))))|(v2_struct_0(X28)|(~v2_pre_topc(X28)|~l1_pre_topc(X28))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_tex_2])])])])])])).
cnf(c_0_43, plain, (k1_tarski(X1)=X2|~r2_hidden(X1,X2)|~v1_zfmisc_1(X2)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34]), c_0_35])).
cnf(c_0_44, negated_conjecture, (r2_hidden(esk6_1(esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
fof(c_0_45, plain, ![X38, X39]:(v2_struct_0(X38)|(~v2_pre_topc(X38)|~l1_pre_topc(X38))|(~m1_subset_1(X39,u1_struct_0(X38))|v2_tex_2(k6_domain_1(u1_struct_0(X38),X39),X38))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t36_tex_2])])])])).
cnf(c_0_46, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_47, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_48, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_40]), c_0_36])).
cnf(c_0_49, plain, (m1_pre_topc(esk4_2(X1,X2),X1)|v1_xboole_0(X2)|v2_struct_0(X1)|~v2_tex_2(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_50, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_51, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_52, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_53, negated_conjecture, (k1_tarski(esk6_1(esk2_0))=esk2_0|~v1_zfmisc_1(esk2_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_54, negated_conjecture, (v2_tex_2(esk2_0,esk1_0)|v1_zfmisc_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_55, plain, ![X34]:((m1_subset_1(esk5_1(X34),k1_zfmisc_1(X34))|v1_xboole_0(X34))&(~v1_xboole_0(esk5_1(X34))|v1_xboole_0(X34))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[rc1_subset_1])])])])])).
fof(c_0_56, plain, ![X23, X24]:(v2_struct_0(X23)|(~v2_pre_topc(X23)|(~v2_tdlat_3(X23)|~l1_pre_topc(X23)))|(~m1_pre_topc(X24,X23)|v2_tdlat_3(X24))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc6_tdlat_3])])])])).
fof(c_0_57, plain, ![X45]:(~l1_pre_topc(X45)|(~v2_tdlat_3(X45)|v2_pre_topc(X45))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_tdlat_3])])).
cnf(c_0_58, plain, (v2_struct_0(X1)|v2_tex_2(k6_domain_1(u1_struct_0(X1),X2),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_59, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),X1)=k1_tarski(X1)|~r2_hidden(X1,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])).
cnf(c_0_60, negated_conjecture, (m1_pre_topc(esk4_2(esk1_0,esk2_0),esk1_0)|~v2_tex_2(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_40]), c_0_50]), c_0_51])]), c_0_36]), c_0_52])).
cnf(c_0_61, negated_conjecture, (k1_tarski(esk6_1(esk2_0))=esk2_0|v2_tex_2(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_62, plain, (v1_tdlat_3(esk4_2(X1,X2))|v1_xboole_0(X2)|v2_struct_0(X1)|~v2_tex_2(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_63, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk5_1(X1))), inference(split_conjunct,[status(thm)],[c_0_55])).
fof(c_0_64, plain, ![X6]:((~v2_struct_0(X6)|(v2_struct_0(X6)|(~v2_pre_topc(X6)|(~v1_tdlat_3(X6)|~v2_tdlat_3(X6))))|~l1_pre_topc(X6))&((v7_struct_0(X6)|(v2_struct_0(X6)|(~v2_pre_topc(X6)|(~v1_tdlat_3(X6)|~v2_tdlat_3(X6))))|~l1_pre_topc(X6))&(v2_pre_topc(X6)|(v2_struct_0(X6)|(~v2_pre_topc(X6)|(~v1_tdlat_3(X6)|~v2_tdlat_3(X6))))|~l1_pre_topc(X6)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_tex_1])])])])).
fof(c_0_65, plain, ![X7]:(~l1_pre_topc(X7)|(~v1_tdlat_3(X7)|v2_pre_topc(X7))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_tdlat_3])])).
cnf(c_0_66, plain, (v2_struct_0(X1)|v2_tdlat_3(X2)|~v2_pre_topc(X1)|~v2_tdlat_3(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_67, plain, (v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_tdlat_3(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_68, negated_conjecture, (v2_tex_2(k1_tarski(X1),esk1_0)|~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_50]), c_0_51])]), c_0_52]), c_0_47])).
cnf(c_0_69, negated_conjecture, (k1_tarski(esk6_1(esk2_0))=esk2_0|m1_pre_topc(esk4_2(esk1_0,esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_70, negated_conjecture, (v1_tdlat_3(esk4_2(esk1_0,esk2_0))|~v2_tex_2(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_40]), c_0_50]), c_0_51])]), c_0_36]), c_0_52])).
fof(c_0_71, plain, ![X18, X19]:(~l1_pre_topc(X18)|(~m1_pre_topc(X19,X18)|l1_pre_topc(X19))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_72, plain, (m1_subset_1(esk5_1(X1),k1_zfmisc_1(X1))|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_73, plain, (r2_hidden(esk6_1(esk5_1(X1)),esk5_1(X1))|v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_63, c_0_37])).
cnf(c_0_74, plain, (v7_struct_0(X1)|v2_struct_0(X1)|~v2_pre_topc(X1)|~v1_tdlat_3(X1)|~v2_tdlat_3(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_75, plain, (v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_76, plain, (v2_tdlat_3(X1)|v2_struct_0(X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_tdlat_3(X2)), inference(csr,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_77, negated_conjecture, (m1_pre_topc(esk4_2(esk1_0,esk2_0),esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_44])]), c_0_60])).
cnf(c_0_78, negated_conjecture, (v2_tdlat_3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_79, negated_conjecture, (k1_tarski(esk6_1(esk2_0))=esk2_0|v1_tdlat_3(esk4_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_70, c_0_61])).
cnf(c_0_80, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
fof(c_0_81, plain, ![X46, X47, X48]:(~r1_tarski(X46,X47)|~r1_tarski(X47,X48)|r1_tarski(X46,X48)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
cnf(c_0_82, plain, (r1_tarski(esk5_1(X1),X1)|v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_25, c_0_72])).
cnf(c_0_83, negated_conjecture, (r2_hidden(esk6_1(esk5_1(esk2_0)),esk5_1(esk2_0))), inference(spm,[status(thm)],[c_0_36, c_0_73])).
cnf(c_0_84, plain, (v1_xboole_0(X2)|v2_struct_0(X1)|~v2_struct_0(esk4_2(X1,X2))|~v2_tex_2(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_85, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_86, plain, (v7_struct_0(X1)|v2_struct_0(X1)|~v1_tdlat_3(X1)|~l1_pre_topc(X1)|~v2_tdlat_3(X1)), inference(csr,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_87, negated_conjecture, (v2_tdlat_3(esk4_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_50]), c_0_78])]), c_0_52])).
cnf(c_0_88, negated_conjecture, (v1_tdlat_3(esk4_2(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_79]), c_0_44])]), c_0_70])).
cnf(c_0_89, negated_conjecture, (l1_pre_topc(esk4_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_77]), c_0_50])])).
cnf(c_0_90, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_91, negated_conjecture, (r1_tarski(esk5_1(esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_36, c_0_82])).
cnf(c_0_92, negated_conjecture, (~v1_xboole_0(esk5_1(esk2_0))), inference(spm,[status(thm)],[c_0_35, c_0_83])).
cnf(c_0_93, plain, (v1_xboole_0(X1)|v2_struct_0(X2)|~r1_tarski(X1,u1_struct_0(X2))|~v2_tex_2(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~v2_struct_0(esk4_2(X2,X1))), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_94, negated_conjecture, (v7_struct_0(esk4_2(esk1_0,esk2_0))|v2_struct_0(esk4_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_88]), c_0_89])])).
cnf(c_0_95, negated_conjecture, (r1_tarski(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_25, c_0_40])).
fof(c_0_96, plain, ![X8, X9, X10, X11, X12]:((~r1_tarski(X8,X9)|(~r2_hidden(X10,X8)|r2_hidden(X10,X9)))&((r2_hidden(esk3_2(X11,X12),X11)|r1_tarski(X11,X12))&(~r2_hidden(esk3_2(X11,X12),X12)|r1_tarski(X11,X12)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_97, negated_conjecture, (r1_tarski(X1,esk2_0)|~r1_tarski(X1,esk5_1(esk2_0))), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_98, negated_conjecture, (r1_tarski(esk5_1(esk5_1(esk2_0)),esk5_1(esk2_0))), inference(spm,[status(thm)],[c_0_92, c_0_82])).
fof(c_0_99, plain, ![X40]:(~l1_pre_topc(X40)|l1_struct_0(X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_100, negated_conjecture, (v7_struct_0(esk4_2(esk1_0,esk2_0))|~v2_tex_2(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_95]), c_0_50]), c_0_51])]), c_0_36]), c_0_52])).
cnf(c_0_101, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_102, negated_conjecture, (r1_tarski(esk5_1(esk5_1(esk2_0)),esk2_0)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
fof(c_0_103, plain, ![X31]:(~v7_struct_0(X31)|~l1_struct_0(X31)|v1_zfmisc_1(u1_struct_0(X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc7_struct_0])])).
cnf(c_0_104, plain, (X1=u1_struct_0(esk4_2(X2,X1))|v1_xboole_0(X1)|v2_struct_0(X2)|~v2_tex_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_105, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_106, negated_conjecture, (k1_tarski(esk6_1(esk2_0))=esk2_0|v7_struct_0(esk4_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_100, c_0_61])).
cnf(c_0_107, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk5_1(esk5_1(esk2_0)))), inference(spm,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_108, negated_conjecture, (esk5_1(esk2_0)=esk2_0|v1_xboole_0(esk5_1(esk2_0))|~v1_zfmisc_1(esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_91]), c_0_36])).
cnf(c_0_109, plain, (v1_zfmisc_1(u1_struct_0(X1))|~v7_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_110, negated_conjecture, (u1_struct_0(esk4_2(esk1_0,esk2_0))=esk2_0|~v2_tex_2(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_40]), c_0_50]), c_0_51])]), c_0_36]), c_0_52])).
cnf(c_0_111, negated_conjecture, (l1_struct_0(esk4_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_105, c_0_89])).
cnf(c_0_112, negated_conjecture, (v7_struct_0(esk4_2(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_106]), c_0_44])]), c_0_100])).
cnf(c_0_113, negated_conjecture, (r2_hidden(esk6_1(esk5_1(esk5_1(esk2_0))),esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_73]), c_0_92])).
cnf(c_0_114, negated_conjecture, (esk5_1(esk2_0)=esk2_0|~v1_zfmisc_1(esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_108]), c_0_36])).
cnf(c_0_115, negated_conjecture, (v1_zfmisc_1(esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_111]), c_0_112])]), c_0_54])).
cnf(c_0_116, negated_conjecture, (k1_tarski(esk6_1(esk5_1(esk5_1(esk2_0))))=esk2_0|~v1_zfmisc_1(esk2_0)), inference(spm,[status(thm)],[c_0_43, c_0_113])).
cnf(c_0_117, negated_conjecture, (esk5_1(esk2_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_114, c_0_115])])).
cnf(c_0_118, negated_conjecture, (~v2_tex_2(esk2_0,esk1_0)|~v1_zfmisc_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_119, negated_conjecture, (k1_tarski(esk6_1(esk2_0))=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_116, c_0_115])]), c_0_117]), c_0_117])).
cnf(c_0_120, negated_conjecture, (~v2_tex_2(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_118, c_0_115])])).
cnf(c_0_121, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_119]), c_0_44])]), c_0_120]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 122
# Proof object clause steps            : 79
# Proof object formula steps           : 43
# Proof object conjectures             : 49
# Proof object clause conjectures      : 46
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 34
# Proof object initial formulas used   : 21
# Proof object generating inferences   : 40
# Proof object simplifying inferences  : 68
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 37
# Processed clauses                    : 328
# ...of these trivial                  : 0
# ...subsumed                          : 100
# ...remaining for further processing  : 227
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 14
# Backward-rewritten                   : 40
# Generated clauses                    : 656
# ...of the previous two non-trivial   : 590
# Contextual simplify-reflections      : 17
# Paramodulations                      : 656
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 136
#    Positive orientable unit clauses  : 23
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 8
#    Non-unit-clauses                  : 105
# Current number of unprocessed clauses: 329
# ...number of literals in the above   : 1704
# Current number of archived formulas  : 0
# Current number of archived clauses   : 91
# Clause-clause subsumption calls (NU) : 3114
# Rec. Clause-clause subsumption calls : 1075
# Non-unit clause-clause subsumptions  : 65
# Unit Clause-clause subsumption calls : 397
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 17
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 13386

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.000 s
# Total time               : 0.020 s
# Maximum resident set size: 28728 pages
