# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t28_zfmisc_1, conjecture, ![X1, X2, X3, X4]:~((r1_tarski(k2_tarski(X1,X2),k2_tarski(X3,X4))&(~(X1=X3)&~(X1=X4)))), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t28_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t75_enumset1)).
fof(t18_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k3_xboole_0(X2,X3))=>r1_tarski(X1,X2)), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t18_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', commutativity_k3_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t28_xboole_1)).
fof(t12_zfmisc_1, axiom, ![X1, X2]:r1_tarski(k1_tarski(X1),k2_tarski(X1,X2)), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t12_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t69_enumset1)).
fof(t68_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k5_enumset1(X1,X2,X3,X4,X5,X6,X7),k1_tarski(X8)), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t68_enumset1)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t98_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t100_xboole_1)).
fof(t107_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X4,X3,X2), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t107_enumset1)).
fof(t92_xboole_1, axiom, ![X1]:k5_xboole_0(X1,X1)=k1_xboole_0, file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t92_xboole_1)).
fof(t5_boole, axiom, ![X1]:k5_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t5_boole)).
fof(t60_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k2_tarski(X6,X7)), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t60_enumset1)).
fof(t66_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_enumset1(X6,X7,X8)), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t66_enumset1)).
fof(t49_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),X3), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t49_xboole_1)).
fof(t16_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3)), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t16_xboole_1)).
fof(t96_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),k5_xboole_0(X1,X2)), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t96_xboole_1)).
fof(t125_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t125_enumset1)).
fof(t69_xboole_1, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>~((r1_tarski(X2,X1)&r1_xboole_0(X2,X1)))), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t69_xboole_1)).
fof(l97_xboole_1, axiom, ![X1, X2]:r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2)), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', l97_xboole_1)).
fof(t136_enumset1, axiom, ![X1, X2, X3]:~((~(X1=X2)&(~(X1=X3)&~(k4_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X1))=k2_tarski(X2,X3))))), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', t136_enumset1)).
fof(fc3_xboole_0, axiom, ![X1, X2]:~(v1_xboole_0(k2_tarski(X1,X2))), file('data/MPTP/proofs/5a01e1d45b254782a697cd801cdb2d9c.E_input', fc3_xboole_0)).
fof(c_0_28, negated_conjecture, ~(![X1, X2, X3, X4]:~((r1_tarski(k2_tarski(X1,X2),k2_tarski(X3,X4))&(~(X1=X3)&~(X1=X4))))), inference(assume_negation,[status(cth)],[t28_zfmisc_1])).
fof(c_0_29, negated_conjecture, (r1_tarski(k2_tarski(esk1_0,esk2_0),k2_tarski(esk3_0,esk4_0))&(esk1_0!=esk3_0&esk1_0!=esk4_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_28])])])])).
fof(c_0_30, plain, ![X13, X14]:k1_enumset1(X13,X13,X14)=k2_tarski(X13,X14), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_31, plain, ![X15, X16, X17]:k2_enumset1(X15,X15,X16,X17)=k1_enumset1(X15,X16,X17), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_32, plain, ![X18, X19, X20, X21]:k3_enumset1(X18,X18,X19,X20,X21)=k2_enumset1(X18,X19,X20,X21), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_33, plain, ![X22, X23, X24, X25, X26]:k4_enumset1(X22,X22,X23,X24,X25,X26)=k3_enumset1(X22,X23,X24,X25,X26), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_34, plain, ![X27, X28, X29, X30, X31, X32]:k5_enumset1(X27,X27,X28,X29,X30,X31,X32)=k4_enumset1(X27,X28,X29,X30,X31,X32), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_35, plain, ![X33, X34, X35, X36, X37, X38, X39]:k6_enumset1(X33,X33,X34,X35,X36,X37,X38,X39)=k5_enumset1(X33,X34,X35,X36,X37,X38,X39), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_36, plain, ![X40, X41, X42]:(~r1_tarski(X40,k3_xboole_0(X41,X42))|r1_tarski(X40,X41)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t18_xboole_1])])).
fof(c_0_37, plain, ![X43, X44]:k3_xboole_0(X43,X44)=k3_xboole_0(X44,X43), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_38, plain, ![X45, X46]:(~r1_tarski(X45,X46)|k3_xboole_0(X45,X46)=X45), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_39, negated_conjecture, (r1_tarski(k2_tarski(esk1_0,esk2_0),k2_tarski(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_40, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_41, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_42, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_44, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_46, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_47, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_48, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_49, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0),k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_40]), c_0_40]), c_0_41]), c_0_41]), c_0_42]), c_0_42]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45])).
cnf(c_0_50, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k3_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_51, negated_conjecture, (k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0),k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
fof(c_0_52, plain, ![X47, X48]:r1_tarski(k1_tarski(X47),k2_tarski(X47,X48)), inference(variable_rename,[status(thm)],[t12_zfmisc_1])).
fof(c_0_53, plain, ![X49]:k2_tarski(X49,X49)=k1_tarski(X49), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_54, plain, ![X50, X51, X52, X53, X54, X55, X56, X57]:k6_enumset1(X50,X51,X52,X53,X54,X55,X56,X57)=k2_xboole_0(k5_enumset1(X50,X51,X52,X53,X54,X55,X56),k1_tarski(X57)), inference(variable_rename,[status(thm)],[t68_enumset1])).
fof(c_0_55, plain, ![X58, X59]:k2_xboole_0(X58,X59)=k5_xboole_0(X58,k4_xboole_0(X59,X58)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_56, plain, ![X60, X61]:k4_xboole_0(X60,X61)=k5_xboole_0(X60,k3_xboole_0(X60,X61)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_57, negated_conjecture, (r1_tarski(X1,k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))|~r1_tarski(X1,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_58, plain, (r1_tarski(k1_tarski(X1),k2_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_59, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
fof(c_0_60, plain, ![X62, X63, X64, X65]:k2_enumset1(X62,X63,X64,X65)=k2_enumset1(X62,X65,X64,X63), inference(variable_rename,[status(thm)],[t107_enumset1])).
cnf(c_0_61, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k5_enumset1(X1,X2,X3,X4,X5,X6,X7),k1_tarski(X8))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_62, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_63, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_64, negated_conjecture, (k3_xboole_0(X1,k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))=X1|~r1_tarski(X1,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_48, c_0_57])).
cnf(c_0_65, plain, (r1_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_59]), c_0_40]), c_0_40]), c_0_41]), c_0_41]), c_0_42]), c_0_42]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45])).
fof(c_0_66, plain, ![X66]:k5_xboole_0(X66,X66)=k1_xboole_0, inference(variable_rename,[status(thm)],[t92_xboole_1])).
fof(c_0_67, plain, ![X67]:k5_xboole_0(X67,k1_xboole_0)=X67, inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_68, plain, (k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X4,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
fof(c_0_69, plain, ![X68, X69, X70, X71, X72, X73, X74]:k5_enumset1(X68,X69,X70,X71,X72,X73,X74)=k2_xboole_0(k3_enumset1(X68,X69,X70,X71,X72),k2_tarski(X73,X74)), inference(variable_rename,[status(thm)],[t60_enumset1])).
fof(c_0_70, plain, ![X75, X76, X77, X78, X79, X80, X81, X82]:k6_enumset1(X75,X76,X77,X78,X79,X80,X81,X82)=k2_xboole_0(k3_enumset1(X75,X76,X77,X78,X79),k1_enumset1(X80,X81,X82)), inference(variable_rename,[status(thm)],[t66_enumset1])).
fof(c_0_71, plain, ![X83, X84, X85]:k3_xboole_0(X83,k4_xboole_0(X84,X85))=k4_xboole_0(k3_xboole_0(X83,X84),X85), inference(variable_rename,[status(thm)],[t49_xboole_1])).
cnf(c_0_72, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k5_xboole_0(k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7),k5_xboole_0(k6_enumset1(X8,X8,X8,X8,X8,X8,X8,X8),k3_xboole_0(k6_enumset1(X8,X8,X8,X8,X8,X8,X8,X8),k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_59]), c_0_40]), c_0_62]), c_0_63]), c_0_41]), c_0_41]), c_0_42]), c_0_42]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_45]), c_0_45])).
cnf(c_0_73, negated_conjecture, (k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_74, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_75, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_76, plain, (k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k6_enumset1(X1,X1,X1,X1,X1,X4,X3,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_42]), c_0_42]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45])).
cnf(c_0_77, plain, (k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k2_tarski(X6,X7))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_78, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_enumset1(X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_79, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_71])).
fof(c_0_80, plain, ![X86, X87, X88]:k3_xboole_0(k3_xboole_0(X86,X87),X88)=k3_xboole_0(X86,k3_xboole_0(X87,X88)), inference(variable_rename,[status(thm)],[t16_xboole_1])).
cnf(c_0_81, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)=k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_74]), c_0_75]), c_0_76])).
cnf(c_0_82, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5),k5_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X6,X6,X7),k3_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X6,X6,X7),k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_40]), c_0_62]), c_0_63]), c_0_41]), c_0_41]), c_0_42]), c_0_42]), c_0_43]), c_0_43]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_45]), c_0_45]), c_0_45])).
cnf(c_0_83, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5),k5_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X6,X7,X8),k3_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X6,X7,X8),k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_62]), c_0_63]), c_0_41]), c_0_41]), c_0_42]), c_0_42]), c_0_43]), c_0_43]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_45]), c_0_45])).
fof(c_0_84, plain, ![X93, X94]:r1_tarski(k4_xboole_0(X93,X94),k5_xboole_0(X93,X94)), inference(variable_rename,[status(thm)],[t96_xboole_1])).
cnf(c_0_85, plain, (k3_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))=k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(k3_xboole_0(X1,X2),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_63]), c_0_63])).
cnf(c_0_86, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_87, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,X1)=k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk1_0,esk4_0,esk3_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_81]), c_0_72])).
cnf(c_0_88, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X6,X7)=k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)), inference(rw,[status(thm)],[c_0_82, c_0_83])).
fof(c_0_89, plain, ![X89, X90, X91, X92]:k2_enumset1(X89,X90,X91,X92)=k2_enumset1(X92,X91,X90,X89), inference(variable_rename,[status(thm)],[t125_enumset1])).
cnf(c_0_90, plain, (r1_tarski(k4_xboole_0(X1,X2),k5_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_91, plain, (k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,k3_xboole_0(X2,X3)))=k3_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))), inference(rw,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_92, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0,X1)=k6_enumset1(esk3_0,esk1_0,esk4_0,esk3_0,esk3_0,esk3_0,esk3_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_88]), c_0_88]), c_0_88])).
cnf(c_0_93, plain, (k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_94, plain, (r1_tarski(k5_xboole_0(X1,k3_xboole_0(X1,X2)),k5_xboole_0(X1,X2))), inference(rw,[status(thm)],[c_0_90, c_0_63])).
cnf(c_0_95, negated_conjecture, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_51]), c_0_74]), c_0_74])).
cnf(c_0_96, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,X1,esk4_0,esk3_0)=k6_enumset1(esk3_0,esk1_0,esk4_0,esk3_0,esk3_0,esk3_0,esk3_0,X1)), inference(spm,[status(thm)],[c_0_76, c_0_92])).
cnf(c_0_97, plain, (k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k6_enumset1(X4,X4,X4,X4,X4,X3,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_42]), c_0_42]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45])).
cnf(c_0_98, negated_conjecture, (r1_tarski(X1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_75]), c_0_75])).
cnf(c_0_99, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,X1,esk4_0,esk4_0,esk3_0)=k6_enumset1(esk3_0,esk1_0,esk4_0,esk3_0,esk3_0,esk3_0,esk3_0,X1)), inference(spm,[status(thm)],[c_0_88, c_0_96])).
cnf(c_0_100, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,X1,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk4_0,esk3_0,esk3_0,esk3_0,esk3_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_96]), c_0_76])).
cnf(c_0_101, plain, (k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5)=k6_enumset1(X1,X1,X1,X1,X4,X3,X2,X5)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_76]), c_0_72])).
cnf(c_0_102, negated_conjecture, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_46, c_0_98])).
cnf(c_0_103, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,X1,esk4_0,esk4_0,esk4_0,esk3_0)=k6_enumset1(esk3_0,esk1_0,esk4_0,esk3_0,esk3_0,esk3_0,esk3_0,X1)), inference(spm,[status(thm)],[c_0_88, c_0_99])).
cnf(c_0_104, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,X1,esk3_0,esk3_0,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk4_0,esk3_0,esk3_0,esk3_0,esk3_0,X1)), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_105, negated_conjecture, (k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_102]), c_0_86])).
cnf(c_0_106, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_88]), c_0_88]), c_0_88]), c_0_88])).
cnf(c_0_107, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,X1,esk4_0,esk4_0,esk4_0,esk4_0,esk3_0)=k6_enumset1(esk3_0,esk1_0,esk4_0,esk3_0,esk3_0,esk3_0,esk3_0,X1)), inference(spm,[status(thm)],[c_0_88, c_0_103])).
cnf(c_0_108, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,X1,esk3_0,esk3_0,esk3_0,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk4_0,esk3_0,esk3_0,esk3_0,esk3_0,X1)), inference(spm,[status(thm)],[c_0_88, c_0_104])).
cnf(c_0_109, negated_conjecture, (k3_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1)))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_105]), c_0_74])).
cnf(c_0_110, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk3_0)=k6_enumset1(esk3_0,esk1_0,esk4_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)), inference(rw,[status(thm)],[c_0_106, c_0_100])).
cnf(c_0_111, negated_conjecture, (k6_enumset1(esk3_0,X1,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk3_0)=k6_enumset1(esk3_0,esk1_0,esk4_0,esk3_0,esk3_0,esk3_0,esk3_0,X1)), inference(spm,[status(thm)],[c_0_88, c_0_107])).
cnf(c_0_112, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,X1,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk4_0,esk3_0,esk3_0,esk3_0,esk3_0,X1)), inference(spm,[status(thm)],[c_0_88, c_0_108])).
cnf(c_0_113, plain, (r1_tarski(k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7),k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_72]), c_0_109]), c_0_75])).
cnf(c_0_114, negated_conjecture, (k6_enumset1(esk3_0,esk1_0,esk4_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)=k6_enumset1(esk3_0,esk1_0,esk4_0,esk3_0,esk3_0,esk3_0,esk3_0,esk1_0)), inference(rw,[status(thm)],[c_0_110, c_0_111])).
cnf(c_0_115, negated_conjecture, (k6_enumset1(esk3_0,X1,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)=k6_enumset1(esk3_0,esk1_0,esk4_0,esk3_0,esk3_0,esk3_0,esk3_0,X1)), inference(spm,[status(thm)],[c_0_88, c_0_112])).
cnf(c_0_116, negated_conjecture, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_50, c_0_98])).
cnf(c_0_117, negated_conjecture, (k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk3_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_81]), c_0_88]), c_0_88]), c_0_88]), c_0_88])).
fof(c_0_118, plain, ![X95, X96]:(v1_xboole_0(X96)|(~r1_tarski(X96,X95)|~r1_xboole_0(X96,X95))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t69_xboole_1])])])).
fof(c_0_119, plain, ![X97, X98]:r1_xboole_0(k3_xboole_0(X97,X98),k5_xboole_0(X97,X98)), inference(variable_rename,[status(thm)],[l97_xboole_1])).
fof(c_0_120, plain, ![X99, X100, X101]:(X99=X100|(X99=X101|k4_xboole_0(k1_enumset1(X99,X100,X101),k1_tarski(X99))=k2_tarski(X100,X101))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t136_enumset1])])])).
cnf(c_0_121, negated_conjecture, (r1_tarski(k6_enumset1(esk3_0,esk1_0,esk4_0,esk3_0,esk3_0,esk3_0,esk3_0,esk1_0),k6_enumset1(esk3_0,X1,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_111]), c_0_88]), c_0_114])).
cnf(c_0_122, negated_conjecture, (k6_enumset1(esk3_0,X1,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk3_0)=k6_enumset1(esk3_0,X1,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_111, c_0_115])).
cnf(c_0_123, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk3_0))), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_124, plain, (v1_xboole_0(X1)|~r1_tarski(X1,X2)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_125, plain, (r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_126, plain, (X1=X2|X1=X3|k4_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X1))=k2_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_127, negated_conjecture, (k3_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_48, c_0_98])).
fof(c_0_128, plain, ![X102, X103]:~v1_xboole_0(k2_tarski(X102,X103)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc3_xboole_0])])).
cnf(c_0_129, negated_conjecture, (r1_tarski(k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0),k6_enumset1(esk3_0,X1,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_121, c_0_115]), c_0_122])).
cnf(c_0_130, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk4_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))), inference(rw,[status(thm)],[c_0_123, c_0_110])).
cnf(c_0_131, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~r1_tarski(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_124, c_0_125])).
cnf(c_0_132, plain, (X1=X3|X1=X2|k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k3_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)))=k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_126, c_0_59]), c_0_40]), c_0_40]), c_0_63]), c_0_41]), c_0_41]), c_0_41]), c_0_41]), c_0_42]), c_0_42]), c_0_42]), c_0_42]), c_0_43]), c_0_43]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_45]), c_0_45])).
cnf(c_0_133, negated_conjecture, (k3_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_86, c_0_127])).
cnf(c_0_134, plain, (k3_xboole_0(k5_xboole_0(X1,X2),k5_xboole_0(X1,k3_xboole_0(X1,X2)))=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_94]), c_0_47])).
cnf(c_0_135, plain, (~v1_xboole_0(k2_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_128])).
cnf(c_0_136, negated_conjecture, (k3_xboole_0(k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0),k6_enumset1(esk3_0,X1,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))=k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_48, c_0_129])).
cnf(c_0_137, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk4_0,esk3_0,esk3_0,esk3_0,esk3_0,esk1_0))), inference(rw,[status(thm)],[c_0_130, c_0_114])).
cnf(c_0_138, plain, (X1=X2|X1=X3|v1_xboole_0(k3_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)))|~r1_tarski(k3_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)),k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_132]), c_0_133]), c_0_133])).
cnf(c_0_139, plain, (k3_xboole_0(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8),k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5))=k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_83]), c_0_109]), c_0_75]), c_0_109]), c_0_75])).
cnf(c_0_140, plain, (~v1_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_135, c_0_40]), c_0_41]), c_0_42]), c_0_43]), c_0_44]), c_0_45])).
cnf(c_0_141, negated_conjecture, (r1_tarski(X1,k6_enumset1(esk3_0,X2,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))|~r1_tarski(X1,k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_50, c_0_136])).
cnf(c_0_142, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(rw,[status(thm)],[c_0_137, c_0_115])).
cnf(c_0_143, plain, (X1=X2|X1=X3|~r1_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X2))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_138, c_0_139]), c_0_139]), c_0_140])).
cnf(c_0_144, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk3_0,X1,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_141, c_0_142])).
cnf(c_0_145, negated_conjecture, (esk1_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_146, negated_conjecture, (esk1_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_147, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144]), c_0_145]), c_0_146]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 148
# Proof object clause steps            : 91
# Proof object formula steps           : 57
# Proof object conjectures             : 47
# Proof object clause conjectures      : 44
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 30
# Proof object initial formulas used   : 28
# Proof object generating inferences   : 38
# Proof object simplifying inferences  : 177
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 28
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 21
# Processed clauses                    : 14030
# ...of these trivial                  : 1347
# ...subsumed                          : 11359
# ...remaining for further processing  : 1324
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 440
# Generated clauses                    : 179357
# ...of the previous two non-trivial   : 135364
# Contextual simplify-reflections      : 0
# Paramodulations                      : 179357
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 863
#    Positive orientable unit clauses  : 482
#    Positive unorientable unit clauses: 56
#    Negative unit clauses             : 180
#    Non-unit-clauses                  : 145
# Current number of unprocessed clauses: 119162
# ...number of literals in the above   : 195805
# Current number of archived formulas  : 0
# Current number of archived clauses   : 470
# Clause-clause subsumption calls (NU) : 12501
# Rec. Clause-clause subsumption calls : 10484
# Non-unit clause-clause subsumptions  : 2847
# Unit Clause-clause subsumption calls : 25487
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 57174
# BW rewrite match successes           : 8008
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3324024

# -------------------------------------------------
# User time                : 1.403 s
# System time              : 0.026 s
# Total time               : 1.429 s
# Maximum resident set size: 96956 pages
