# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/5ab934fbb4b241e79a9f35af213fac60.E_input', d3_tarski)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/5ab934fbb4b241e79a9f35af213fac60.E_input', t2_xboole_1)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/5ab934fbb4b241e79a9f35af213fac60.E_input', d4_relat_1)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/5ab934fbb4b241e79a9f35af213fac60.E_input', l54_zfmisc_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/5ab934fbb4b241e79a9f35af213fac60.E_input', d5_relat_1)).
fof(t195_relat_1, conjecture, ![X1, X2]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&~((k1_relat_1(k2_zfmisc_1(X1,X2))=X1&k2_relat_1(k2_zfmisc_1(X1,X2))=X2))))), file('data/MPTP/proofs/5ab934fbb4b241e79a9f35af213fac60.E_input', t195_relat_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/5ab934fbb4b241e79a9f35af213fac60.E_input', d10_xboole_0)).
fof(c_0_7, plain, ![X36, X37, X38, X39, X40]:((~r1_tarski(X36,X37)|(~r2_hidden(X38,X36)|r2_hidden(X38,X37)))&((r2_hidden(esk9_2(X39,X40),X39)|r1_tarski(X39,X40))&(~r2_hidden(esk9_2(X39,X40),X40)|r1_tarski(X39,X40)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_8, plain, ![X22]:r1_tarski(k1_xboole_0,X22), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_9, plain, ![X23, X24, X25, X27, X28, X29, X30, X32]:(((~r2_hidden(X25,X24)|r2_hidden(k4_tarski(X25,esk6_3(X23,X24,X25)),X23)|X24!=k1_relat_1(X23))&(~r2_hidden(k4_tarski(X27,X28),X23)|r2_hidden(X27,X24)|X24!=k1_relat_1(X23)))&((~r2_hidden(esk7_2(X29,X30),X30)|~r2_hidden(k4_tarski(esk7_2(X29,X30),X32),X29)|X30=k1_relat_1(X29))&(r2_hidden(esk7_2(X29,X30),X30)|r2_hidden(k4_tarski(esk7_2(X29,X30),esk8_2(X29,X30)),X29)|X30=k1_relat_1(X29)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
cnf(c_0_10, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (r2_hidden(k4_tarski(X1,esk6_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_13, plain, ![X18, X19, X20, X21]:(((r2_hidden(X18,X20)|~r2_hidden(k4_tarski(X18,X19),k2_zfmisc_1(X20,X21)))&(r2_hidden(X19,X21)|~r2_hidden(k4_tarski(X18,X19),k2_zfmisc_1(X20,X21))))&(~r2_hidden(X18,X20)|~r2_hidden(X19,X21)|r2_hidden(k4_tarski(X18,X19),k2_zfmisc_1(X20,X21)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
cnf(c_0_14, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15, plain, (r2_hidden(k4_tarski(X1,esk6_3(X2,k1_relat_1(X2),X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_12])).
fof(c_0_16, plain, ![X7, X8, X9, X11, X12, X13, X14, X16]:(((~r2_hidden(X9,X8)|r2_hidden(k4_tarski(esk3_3(X7,X8,X9),X9),X7)|X8!=k2_relat_1(X7))&(~r2_hidden(k4_tarski(X12,X11),X7)|r2_hidden(X11,X8)|X8!=k2_relat_1(X7)))&((~r2_hidden(esk4_2(X13,X14),X14)|~r2_hidden(k4_tarski(X16,esk4_2(X13,X14)),X13)|X14=k2_relat_1(X13))&(r2_hidden(esk4_2(X13,X14),X14)|r2_hidden(k4_tarski(esk5_2(X13,X14),esk4_2(X13,X14)),X13)|X14=k2_relat_1(X13)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_17, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, plain, (r2_hidden(k4_tarski(X1,esk6_3(k1_xboole_0,k1_relat_1(k1_xboole_0),X1)),X2)|~r2_hidden(X1,k1_relat_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_14, c_0_15])).
fof(c_0_19, negated_conjecture, ~(![X1, X2]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&~((k1_relat_1(k2_zfmisc_1(X1,X2))=X1&k2_relat_1(k2_zfmisc_1(X1,X2))=X2)))))), inference(assume_negation,[status(cth)],[t195_relat_1])).
cnf(c_0_20, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,X2))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (r2_hidden(esk4_2(X1,X2),X2)|r2_hidden(k4_tarski(esk5_2(X1,X2),esk4_2(X1,X2)),X1)|X2=k2_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, plain, ![X34, X35]:(((r1_tarski(X34,X35)|X34!=X35)&(r1_tarski(X35,X34)|X34!=X35))&(~r1_tarski(X34,X35)|~r1_tarski(X35,X34)|X34=X35)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_23, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_relat_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_24, plain, (r2_hidden(esk9_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_25, negated_conjecture, (esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(k1_relat_1(k2_zfmisc_1(esk1_0,esk2_0))!=esk1_0|k2_relat_1(k2_zfmisc_1(esk1_0,esk2_0))!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_19])])])])).
cnf(c_0_26, plain, (X1=k2_relat_1(k2_zfmisc_1(X2,X3))|r2_hidden(esk4_2(k2_zfmisc_1(X2,X3),X1),X1)|r2_hidden(esk4_2(k2_zfmisc_1(X2,X3),X1),X3)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27, plain, (r2_hidden(esk7_2(X1,X2),X2)|r2_hidden(k4_tarski(esk7_2(X1,X2),esk8_2(X1,X2)),X1)|X2=k1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_28, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, plain, (r1_tarski(X1,X2)|~r2_hidden(esk9_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_30, plain, (r1_tarski(k1_relat_1(k1_xboole_0),X1)|r2_hidden(esk9_2(k1_relat_1(k1_xboole_0),X1),X2)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_31, negated_conjecture, (k1_relat_1(k2_zfmisc_1(esk1_0,esk2_0))!=esk1_0|k2_relat_1(k2_zfmisc_1(esk1_0,esk2_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, plain, (k2_relat_1(k2_zfmisc_1(X1,X2))=X2|r2_hidden(esk4_2(k2_zfmisc_1(X1,X2),X2),X2)), inference(ef,[status(thm)],[c_0_26])).
cnf(c_0_33, plain, (X1=k1_relat_1(k2_zfmisc_1(X2,X3))|r2_hidden(esk7_2(k2_zfmisc_1(X2,X3),X1),X1)|r2_hidden(esk7_2(k2_zfmisc_1(X2,X3),X1),X2)), inference(spm,[status(thm)],[c_0_17, c_0_27])).
cnf(c_0_34, plain, (r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_35, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_28, c_0_11])).
cnf(c_0_36, plain, (r1_tarski(k1_relat_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_37, plain, (X2=k2_relat_1(X1)|~r2_hidden(esk4_2(X1,X2),X2)|~r2_hidden(k4_tarski(X3,esk4_2(X1,X2)),X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_38, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_39, negated_conjecture, (r2_hidden(esk4_2(k2_zfmisc_1(esk1_0,esk2_0),esk2_0),esk2_0)|k1_relat_1(k2_zfmisc_1(esk1_0,esk2_0))!=esk1_0), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_40, plain, (k1_relat_1(k2_zfmisc_1(X1,X2))=X1|r2_hidden(esk7_2(k2_zfmisc_1(X1,X2),X1),X1)), inference(ef,[status(thm)],[c_0_33])).
cnf(c_0_41, plain, (X1=k1_relat_1(X2)|r2_hidden(esk7_2(X2,X1),X1)|r2_hidden(esk7_2(X2,X1),X3)|X3!=k1_relat_1(X2)), inference(spm,[status(thm)],[c_0_34, c_0_27])).
cnf(c_0_42, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_43, plain, (X1=k2_relat_1(k2_zfmisc_1(X2,X3))|~r2_hidden(esk4_2(k2_zfmisc_1(X2,X3),X1),X1)|~r2_hidden(esk4_2(k2_zfmisc_1(X2,X3),X1),X3)|~r2_hidden(X4,X2)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_44, negated_conjecture, (r2_hidden(esk7_2(k2_zfmisc_1(esk1_0,esk2_0),esk1_0),esk1_0)|r2_hidden(esk4_2(k2_zfmisc_1(esk1_0,esk2_0),esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_45, plain, (X1=k1_xboole_0|r2_hidden(esk7_2(k1_xboole_0,X1),X2)|r2_hidden(esk7_2(k1_xboole_0,X1),X1)|X2!=k1_xboole_0), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_46, negated_conjecture, (k2_relat_1(k2_zfmisc_1(esk1_0,esk2_0))=esk2_0|r2_hidden(esk7_2(k2_zfmisc_1(esk1_0,esk2_0),esk1_0),esk1_0)|~r2_hidden(X1,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_32])).
cnf(c_0_47, plain, (X1=k1_xboole_0|r2_hidden(esk7_2(k1_xboole_0,X1),X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_45]), c_0_14])).
cnf(c_0_48, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_49, plain, (X2=k1_relat_1(X1)|~r2_hidden(esk7_2(X1,X2),X2)|~r2_hidden(k4_tarski(esk7_2(X1,X2),X3),X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_50, negated_conjecture, (k2_relat_1(k2_zfmisc_1(esk1_0,esk2_0))=esk2_0|r2_hidden(esk7_2(k2_zfmisc_1(esk1_0,esk2_0),esk1_0),esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])).
cnf(c_0_51, plain, (X1=k1_relat_1(k2_zfmisc_1(X2,X3))|~r2_hidden(esk7_2(k2_zfmisc_1(X2,X3),X1),X1)|~r2_hidden(esk7_2(k2_zfmisc_1(X2,X3),X1),X2)|~r2_hidden(X4,X3)), inference(spm,[status(thm)],[c_0_49, c_0_38])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk7_2(k2_zfmisc_1(esk1_0,esk2_0),esk1_0),esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_50]), c_0_40])).
cnf(c_0_53, negated_conjecture, (k1_relat_1(k2_zfmisc_1(esk1_0,esk2_0))=esk1_0|~r2_hidden(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_52])])).
cnf(c_0_54, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_55, negated_conjecture, (k1_relat_1(k2_zfmisc_1(esk1_0,esk2_0))=esk1_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_47]), c_0_54])).
cnf(c_0_56, negated_conjecture, (r2_hidden(esk4_2(k2_zfmisc_1(esk1_0,esk2_0),esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_55])])).
cnf(c_0_57, negated_conjecture, (k2_relat_1(k2_zfmisc_1(esk1_0,esk2_0))!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_55])])).
cnf(c_0_58, negated_conjecture, (~r2_hidden(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_56]), c_0_56])]), c_0_57])).
cnf(c_0_59, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_52, c_0_58]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 60
# Proof object clause steps            : 45
# Proof object formula steps           : 15
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 15
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 21
# Processed clauses                    : 298
# ...of these trivial                  : 3
# ...subsumed                          : 105
# ...remaining for further processing  : 190
# Other redundant clauses eliminated   : 87
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 7
# Backward-rewritten                   : 64
# Generated clauses                    : 4282
# ...of the previous two non-trivial   : 4129
# Contextual simplify-reflections      : 8
# Paramodulations                      : 4165
# Factorizations                       : 8
# Equation resolutions                 : 107
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 95
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 85
# Current number of unprocessed clauses: 3800
# ...number of literals in the above   : 21208
# Current number of archived formulas  : 0
# Current number of archived clauses   : 93
# Clause-clause subsumption calls (NU) : 8268
# Rec. Clause-clause subsumption calls : 2580
# Non-unit clause-clause subsumptions  : 120
# Unit Clause-clause subsumption calls : 259
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12
# BW rewrite match successes           : 7
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 110732

# -------------------------------------------------
# User time                : 0.062 s
# System time              : 0.003 s
# Total time               : 0.066 s
# Maximum resident set size: 28544 pages
