# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.007 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t125_enumset1, conjecture, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1), file('data/MPTP/proofs/5abb52384cc1474fa343fb8eb46103a8.E_input', t125_enumset1)).
fof(t85_enumset1, axiom, ![X1, X2, X3, X4]:k5_enumset1(X1,X1,X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/5abb52384cc1474fa343fb8eb46103a8.E_input', t85_enumset1)).
fof(t117_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X3,X2,X4,X1), file('data/MPTP/proofs/5abb52384cc1474fa343fb8eb46103a8.E_input', t117_enumset1)).
fof(t48_enumset1, axiom, ![X1, X2, X3, X4, X5]:k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_tarski(X1,X2),k1_enumset1(X3,X4,X5)), file('data/MPTP/proofs/5abb52384cc1474fa343fb8eb46103a8.E_input', t48_enumset1)).
fof(t77_enumset1, axiom, ![X1, X2]:k2_enumset1(X1,X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/5abb52384cc1474fa343fb8eb46103a8.E_input', t77_enumset1)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/5abb52384cc1474fa343fb8eb46103a8.E_input', t98_xboole_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/5abb52384cc1474fa343fb8eb46103a8.E_input', t71_enumset1)).
fof(t80_enumset1, axiom, ![X1, X2, X3, X4, X5]:k5_enumset1(X1,X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/5abb52384cc1474fa343fb8eb46103a8.E_input', t80_enumset1)).
fof(t102_enumset1, axiom, ![X1, X2, X3]:k1_enumset1(X1,X2,X3)=k1_enumset1(X3,X2,X1), file('data/MPTP/proofs/5abb52384cc1474fa343fb8eb46103a8.E_input', t102_enumset1)).
fof(t99_enumset1, axiom, ![X1, X2, X3]:k1_enumset1(X1,X2,X3)=k1_enumset1(X2,X1,X3), file('data/MPTP/proofs/5abb52384cc1474fa343fb8eb46103a8.E_input', t99_enumset1)).
fof(t98_enumset1, axiom, ![X1, X2, X3]:k1_enumset1(X1,X2,X3)=k1_enumset1(X1,X3,X2), file('data/MPTP/proofs/5abb52384cc1474fa343fb8eb46103a8.E_input', t98_enumset1)).
fof(c_0_11, negated_conjecture, ~(![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1)), inference(assume_negation,[status(cth)],[t125_enumset1])).
fof(c_0_12, negated_conjecture, k2_enumset1(esk1_0,esk2_0,esk3_0,esk4_0)!=k2_enumset1(esk4_0,esk3_0,esk2_0,esk1_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
fof(c_0_13, plain, ![X40, X41, X42, X43]:k5_enumset1(X40,X40,X40,X40,X41,X42,X43)=k2_enumset1(X40,X41,X42,X43), inference(variable_rename,[status(thm)],[t85_enumset1])).
fof(c_0_14, plain, ![X33, X34, X35, X36]:k2_enumset1(X33,X34,X35,X36)=k2_enumset1(X35,X34,X36,X33), inference(variable_rename,[status(thm)],[t117_enumset1])).
fof(c_0_15, plain, ![X25, X26, X27, X28, X29]:k3_enumset1(X25,X26,X27,X28,X29)=k2_xboole_0(k2_tarski(X25,X26),k1_enumset1(X27,X28,X29)), inference(variable_rename,[status(thm)],[t48_enumset1])).
fof(c_0_16, plain, ![X23, X24]:k2_enumset1(X23,X23,X23,X24)=k2_tarski(X23,X24), inference(variable_rename,[status(thm)],[t77_enumset1])).
fof(c_0_17, plain, ![X16, X17]:k2_xboole_0(X16,X17)=k5_xboole_0(X16,k4_xboole_0(X17,X16)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_18, plain, ![X13, X14, X15]:k2_enumset1(X13,X13,X14,X15)=k1_enumset1(X13,X14,X15), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_19, plain, ![X18, X19, X20, X21, X22]:k5_enumset1(X18,X18,X18,X19,X20,X21,X22)=k3_enumset1(X18,X19,X20,X21,X22), inference(variable_rename,[status(thm)],[t80_enumset1])).
fof(c_0_20, plain, ![X37, X38, X39]:k1_enumset1(X37,X38,X39)=k1_enumset1(X39,X38,X37), inference(variable_rename,[status(thm)],[t102_enumset1])).
cnf(c_0_21, negated_conjecture, (k2_enumset1(esk1_0,esk2_0,esk3_0,esk4_0)!=k2_enumset1(esk4_0,esk3_0,esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22, plain, (k5_enumset1(X1,X1,X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23, plain, (k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X3,X2,X4,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_24, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_tarski(X1,X2),k1_enumset1(X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25, plain, (k2_enumset1(X1,X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, plain, (k5_enumset1(X1,X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (k1_enumset1(X1,X2,X3)=k1_enumset1(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_30, plain, ![X30, X31, X32]:k1_enumset1(X30,X31,X32)=k1_enumset1(X31,X30,X32), inference(variable_rename,[status(thm)],[t99_enumset1])).
cnf(c_0_31, negated_conjecture, (k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk3_0,esk2_0,esk1_0)!=k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_22]), c_0_22])).
cnf(c_0_32, plain, (k5_enumset1(X1,X1,X1,X1,X2,X3,X4)=k5_enumset1(X3,X3,X3,X3,X2,X4,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_22]), c_0_22])).
cnf(c_0_33, plain, (k5_enumset1(X1,X1,X1,X2,X3,X4,X5)=k5_xboole_0(k5_enumset1(X1,X1,X1,X1,X1,X1,X2),k4_xboole_0(k5_enumset1(X3,X3,X3,X3,X3,X4,X5),k5_enumset1(X1,X1,X1,X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25]), c_0_26]), c_0_27]), c_0_22]), c_0_22]), c_0_22]), c_0_28])).
cnf(c_0_34, plain, (k5_enumset1(X1,X1,X1,X1,X1,X2,X3)=k5_enumset1(X3,X3,X3,X3,X3,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_27]), c_0_27]), c_0_22]), c_0_22])).
cnf(c_0_35, plain, (k1_enumset1(X1,X2,X3)=k1_enumset1(X2,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_36, plain, ![X10, X11, X12]:k1_enumset1(X10,X11,X12)=k1_enumset1(X10,X12,X11), inference(variable_rename,[status(thm)],[t98_enumset1])).
cnf(c_0_37, negated_conjecture, (k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk2_0,esk1_0,esk3_0)!=k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk3_0,esk2_0,esk1_0)), inference(rw,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_38, plain, (k5_enumset1(X1,X1,X1,X2,X3,X4,X5)=k5_enumset1(X1,X1,X1,X2,X5,X4,X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_33])).
cnf(c_0_39, plain, (k5_enumset1(X1,X1,X1,X1,X1,X2,X3)=k5_enumset1(X2,X2,X2,X2,X2,X1,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_27]), c_0_27]), c_0_22]), c_0_22])).
cnf(c_0_40, plain, (k1_enumset1(X1,X2,X3)=k1_enumset1(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_41, negated_conjecture, (k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk3_0,esk1_0,esk2_0)!=k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_38]), c_0_38])).
cnf(c_0_42, plain, (k5_enumset1(X1,X1,X1,X2,X3,X4,X5)=k5_enumset1(X1,X1,X1,X2,X4,X3,X5)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_39]), c_0_33])).
cnf(c_0_43, plain, (k5_enumset1(X1,X1,X1,X1,X1,X2,X3)=k5_enumset1(X1,X1,X1,X1,X1,X3,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_27]), c_0_27]), c_0_22]), c_0_22])).
cnf(c_0_44, negated_conjecture, (k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk1_0,esk2_0,esk3_0)!=k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk1_0,esk3_0,esk2_0)), inference(rw,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_45, plain, (k5_enumset1(X1,X1,X1,X2,X3,X4,X5)=k5_enumset1(X1,X1,X1,X2,X3,X5,X4)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_43]), c_0_33])).
cnf(c_0_46, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_45])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 47
# Proof object clause steps            : 24
# Proof object formula steps           : 23
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 32
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 11
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 6
# Processed clauses                    : 80
# ...of these trivial                  : 0
# ...subsumed                          : 48
# ...remaining for further processing  : 32
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 4
# Generated clauses                    : 892
# ...of the previous two non-trivial   : 788
# Contextual simplify-reflections      : 0
# Paramodulations                      : 892
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 21
#    Positive orientable unit clauses  : 1
#    Positive unorientable unit clauses: 20
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 720
# ...number of literals in the above   : 720
# Current number of archived formulas  : 0
# Current number of archived clauses   : 16
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 231
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 606
# BW rewrite match successes           : 252
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2918

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28572 pages
