# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t39_tops_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k2_pre_topc(X1,X2))<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>~((v3_pre_topc(X4,X1)&(r2_hidden(X3,X4)&r1_xboole_0(X2,X4))))))))), file('data/MPTP/proofs/5ae9ef422adc43d1b993fe37b6b1f0b8.E_input', t39_tops_1)).
fof(d13_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k2_pre_topc(X1,X2)<=>![X4]:(r2_hidden(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))=>~((v3_pre_topc(X5,X1)&(r2_hidden(X4,X5)&r1_xboole_0(X2,X5))))))))))), file('data/MPTP/proofs/5ae9ef422adc43d1b993fe37b6b1f0b8.E_input', d13_pre_topc)).
fof(dt_k2_pre_topc, axiom, ![X1, X2]:((l1_pre_topc(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k2_pre_topc(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/5ae9ef422adc43d1b993fe37b6b1f0b8.E_input', dt_k2_pre_topc)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/5ae9ef422adc43d1b993fe37b6b1f0b8.E_input', t2_subset)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/5ae9ef422adc43d1b993fe37b6b1f0b8.E_input', fc2_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/5ae9ef422adc43d1b993fe37b6b1f0b8.E_input', dt_l1_pre_topc)).
fof(c_0_6, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k2_pre_topc(X1,X2))<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>~((v3_pre_topc(X4,X1)&(r2_hidden(X3,X4)&r1_xboole_0(X2,X4)))))))))), inference(assume_negation,[status(cth)],[t39_tops_1])).
fof(c_0_7, plain, ![X11, X12, X13, X14, X15, X19]:(((~r2_hidden(X14,X13)|(~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X11)))|(~v3_pre_topc(X15,X11)|(~r2_hidden(X14,X15)|~r1_xboole_0(X12,X15))))|~r2_hidden(X14,u1_struct_0(X11))|X13!=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))&((m1_subset_1(esk5_4(X11,X12,X13,X14),k1_zfmisc_1(u1_struct_0(X11)))|r2_hidden(X14,X13)|~r2_hidden(X14,u1_struct_0(X11))|X13!=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))&((v3_pre_topc(esk5_4(X11,X12,X13,X14),X11)|r2_hidden(X14,X13)|~r2_hidden(X14,u1_struct_0(X11))|X13!=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))&((r2_hidden(X14,esk5_4(X11,X12,X13,X14))|r2_hidden(X14,X13)|~r2_hidden(X14,u1_struct_0(X11))|X13!=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))&(r1_xboole_0(X12,esk5_4(X11,X12,X13,X14))|r2_hidden(X14,X13)|~r2_hidden(X14,u1_struct_0(X11))|X13!=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))))))&((r2_hidden(esk6_3(X11,X12,X13),u1_struct_0(X11))|X13=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))&(((m1_subset_1(esk7_3(X11,X12,X13),k1_zfmisc_1(u1_struct_0(X11)))|~r2_hidden(esk6_3(X11,X12,X13),X13)|X13=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))&((v3_pre_topc(esk7_3(X11,X12,X13),X11)|~r2_hidden(esk6_3(X11,X12,X13),X13)|X13=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))&((r2_hidden(esk6_3(X11,X12,X13),esk7_3(X11,X12,X13))|~r2_hidden(esk6_3(X11,X12,X13),X13)|X13=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))&(r1_xboole_0(X12,esk7_3(X11,X12,X13))|~r2_hidden(esk6_3(X11,X12,X13),X13)|X13=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11)))))&(r2_hidden(esk6_3(X11,X12,X13),X13)|(~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X11)))|(~v3_pre_topc(X19,X11)|(~r2_hidden(esk6_3(X11,X12,X13),X19)|~r1_xboole_0(X12,X19))))|X13=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d13_pre_topc])])])])])).
fof(c_0_8, plain, ![X20, X21]:(~l1_pre_topc(X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|m1_subset_1(k2_pre_topc(X20,X21),k1_zfmisc_1(u1_struct_0(X20)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_pre_topc])])).
fof(c_0_9, negated_conjecture, ![X10]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(((m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0)))&((v3_pre_topc(esk4_0,esk1_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0)))&((r2_hidden(esk3_0,esk4_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0)))&(r1_xboole_0(esk2_0,esk4_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))))))&(r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|(~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(esk1_0)))|(~v3_pre_topc(X10,esk1_0)|(~r2_hidden(esk3_0,X10)|~r1_xboole_0(esk2_0,X10))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_6])])])])])])).
cnf(c_0_10, plain, (r1_xboole_0(X1,esk5_4(X2,X1,X3,X4))|r2_hidden(X4,X3)|~r2_hidden(X4,u1_struct_0(X2))|X3!=k2_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (m1_subset_1(k2_pre_topc(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_12, plain, ![X22, X23]:(~m1_subset_1(X22,X23)|(v1_xboole_0(X23)|r2_hidden(X22,X23))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_13, negated_conjecture, (r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v3_pre_topc(X1,esk1_0)|~r2_hidden(esk3_0,X1)|~r1_xboole_0(esk2_0,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (r1_xboole_0(X1,esk5_4(X2,X1,k2_pre_topc(X2,X1),X3))|r2_hidden(X3,k2_pre_topc(X2,X1))|~r2_hidden(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_10]), c_0_11])).
cnf(c_0_15, plain, (v3_pre_topc(esk5_4(X1,X2,X3,X4),X1)|r2_hidden(X4,X3)|~r2_hidden(X4,u1_struct_0(X1))|X3!=k2_pre_topc(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_16, plain, ![X24]:(v2_struct_0(X24)|~l1_struct_0(X24)|~v1_xboole_0(u1_struct_0(X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_17, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_19, negated_conjecture, (r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|r2_hidden(X1,k2_pre_topc(X2,esk2_0))|~v3_pre_topc(esk5_4(X2,esk2_0,k2_pre_topc(X2,esk2_0),X1),esk1_0)|~r2_hidden(esk3_0,esk5_4(X2,esk2_0,k2_pre_topc(X2,esk2_0),X1))|~r2_hidden(X1,u1_struct_0(X2))|~m1_subset_1(esk5_4(X2,esk2_0,k2_pre_topc(X2,esk2_0),X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_20, plain, (v3_pre_topc(esk5_4(X1,X2,k2_pre_topc(X1,X2),X3),X1)|r2_hidden(X3,k2_pre_topc(X1,X2))|~r2_hidden(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_15]), c_0_11])).
cnf(c_0_21, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_22, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_23, plain, (m1_subset_1(esk5_4(X1,X2,X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|r2_hidden(X4,X3)|~r2_hidden(X4,u1_struct_0(X1))|X3!=k2_pre_topc(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_24, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk3_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_26, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_27, plain, ![X25]:(~l1_pre_topc(X25)|l1_struct_0(X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_28, negated_conjecture, (r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|r2_hidden(X1,k2_pre_topc(esk1_0,esk2_0))|~r2_hidden(esk3_0,esk5_4(esk1_0,esk2_0,k2_pre_topc(esk1_0,esk2_0),X1))|~r2_hidden(X1,u1_struct_0(esk1_0))|~m1_subset_1(esk5_4(esk1_0,esk2_0,k2_pre_topc(esk1_0,esk2_0),X1),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_22])])).
cnf(c_0_29, plain, (r2_hidden(X1,k2_pre_topc(X2,X3))|m1_subset_1(esk5_4(X2,X3,k2_pre_topc(X2,X3),X1),k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_23]), c_0_11])).
cnf(c_0_30, plain, (r2_hidden(X1,esk5_4(X2,X3,X4,X1))|r2_hidden(X1,X4)|~r2_hidden(X1,u1_struct_0(X2))|X4!=k2_pre_topc(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_31, negated_conjecture, (r2_hidden(esk3_0,u1_struct_0(esk1_0))|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26])).
cnf(c_0_32, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X4)))|~v3_pre_topc(X3,X4)|~r2_hidden(X1,X3)|~r1_xboole_0(X5,X3)|~r2_hidden(X1,u1_struct_0(X4))|X2!=k2_pre_topc(X4,X5)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X4)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4)))|~l1_pre_topc(X4)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_34, negated_conjecture, (r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|r2_hidden(X1,k2_pre_topc(esk1_0,esk2_0))|~r2_hidden(esk3_0,esk5_4(esk1_0,esk2_0,k2_pre_topc(esk1_0,esk2_0),X1))|~r2_hidden(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_21]), c_0_22])])).
cnf(c_0_35, plain, (r2_hidden(X1,esk5_4(X2,X3,k2_pre_topc(X2,X3),X1))|r2_hidden(X1,k2_pre_topc(X2,X3))|~r2_hidden(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_30]), c_0_11])).
cnf(c_0_36, negated_conjecture, (r2_hidden(esk3_0,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_22])])).
cnf(c_0_37, plain, (~r1_xboole_0(X1,X2)|~v3_pre_topc(X2,X3)|~r2_hidden(X4,k2_pre_topc(X3,X1))|~r2_hidden(X4,u1_struct_0(X3))|~r2_hidden(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_33]), c_0_11])).
cnf(c_0_38, negated_conjecture, (r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_21]), c_0_22])])).
cnf(c_0_39, negated_conjecture, (r1_xboole_0(esk2_0,esk4_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_40, negated_conjecture, (v3_pre_topc(esk4_0,esk1_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_41, negated_conjecture, (r2_hidden(esk3_0,esk4_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_42, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_43, negated_conjecture, (~r1_xboole_0(esk2_0,X1)|~v3_pre_topc(X1,esk1_0)|~r2_hidden(esk3_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_36]), c_0_21]), c_0_22])])).
cnf(c_0_44, negated_conjecture, (r1_xboole_0(esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_38])])).
cnf(c_0_45, negated_conjecture, (v3_pre_topc(esk4_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_38])])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_38])])).
cnf(c_0_47, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_38])])).
cnf(c_0_48, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_46]), c_0_47])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 49
# Proof object clause steps            : 36
# Proof object formula steps           : 13
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 39
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 25
# Processed clauses                    : 76
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 76
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 11
# Generated clauses                    : 30
# ...of the previous two non-trivial   : 32
# Contextual simplify-reflections      : 5
# Paramodulations                      : 25
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 33
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 22
# Current number of unprocessed clauses: 4
# ...number of literals in the above   : 26
# Current number of archived formulas  : 0
# Current number of archived clauses   : 38
# Clause-clause subsumption calls (NU) : 878
# Rec. Clause-clause subsumption calls : 30
# Non-unit clause-clause subsumptions  : 7
# Unit Clause-clause subsumption calls : 32
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4012

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28616 pages
