# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d21_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X3=k15_lattice3(X1,X2)<=>(r4_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r4_lattice3(X1,X4,X2)=>r1_lattices(X1,X3,X4)))))))), file('data/MPTP/proofs/5b123d44fab94d09a814a31003bec19d.E_input', d21_lattice3)).
fof(t38_lattice3, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r2_hidden(X2,X3)=>(r3_lattices(X1,X2,k15_lattice3(X1,X3))&r3_lattices(X1,k16_lattice3(X1,X3),X2))))), file('data/MPTP/proofs/5b123d44fab94d09a814a31003bec19d.E_input', t38_lattice3)).
fof(dt_k15_lattice3, axiom, ![X1, X2]:((~(v2_struct_0(X1))&l3_lattices(X1))=>m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/5b123d44fab94d09a814a31003bec19d.E_input', dt_k15_lattice3)).
fof(redefinition_r3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_lattices(X1,X2,X3)<=>r1_lattices(X1,X2,X3))), file('data/MPTP/proofs/5b123d44fab94d09a814a31003bec19d.E_input', redefinition_r3_lattices)).
fof(d22_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:k16_lattice3(X1,X2)=k15_lattice3(X1,a_2_1_lattice3(X1,X2))), file('data/MPTP/proofs/5b123d44fab94d09a814a31003bec19d.E_input', d22_lattice3)).
fof(fraenkel_a_2_1_lattice3, axiom, ![X1, X2, X3]:((~(v2_struct_0(X2))&l3_lattices(X2))=>(r2_hidden(X1,a_2_1_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&r3_lattice3(X2,X4,X3))))), file('data/MPTP/proofs/5b123d44fab94d09a814a31003bec19d.E_input', fraenkel_a_2_1_lattice3)).
fof(d17_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r4_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X4,X2)))))), file('data/MPTP/proofs/5b123d44fab94d09a814a31003bec19d.E_input', d17_lattice3)).
fof(d16_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r3_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X2,X4)))))), file('data/MPTP/proofs/5b123d44fab94d09a814a31003bec19d.E_input', d16_lattice3)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/5b123d44fab94d09a814a31003bec19d.E_input', cc1_lattices)).
fof(c_0_9, plain, ![X16, X17, X18, X19]:(((r4_lattice3(X16,X18,X17)|X18!=k15_lattice3(X16,X17)|~m1_subset_1(X18,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v4_lattice3(X16)|~l3_lattices(X16))))|(v2_struct_0(X16)|~l3_lattices(X16)))&(~m1_subset_1(X19,u1_struct_0(X16))|(~r4_lattice3(X16,X19,X17)|r1_lattices(X16,X18,X19))|X18!=k15_lattice3(X16,X17)|~m1_subset_1(X18,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v4_lattice3(X16)|~l3_lattices(X16))))|(v2_struct_0(X16)|~l3_lattices(X16))))&((m1_subset_1(esk5_3(X16,X17,X18),u1_struct_0(X16))|~r4_lattice3(X16,X18,X17)|X18=k15_lattice3(X16,X17)|~m1_subset_1(X18,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v4_lattice3(X16)|~l3_lattices(X16))))|(v2_struct_0(X16)|~l3_lattices(X16)))&((r4_lattice3(X16,esk5_3(X16,X17,X18),X17)|~r4_lattice3(X16,X18,X17)|X18=k15_lattice3(X16,X17)|~m1_subset_1(X18,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v4_lattice3(X16)|~l3_lattices(X16))))|(v2_struct_0(X16)|~l3_lattices(X16)))&(~r1_lattices(X16,X18,esk5_3(X16,X17,X18))|~r4_lattice3(X16,X18,X17)|X18=k15_lattice3(X16,X17)|~m1_subset_1(X18,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v4_lattice3(X16)|~l3_lattices(X16))))|(v2_struct_0(X16)|~l3_lattices(X16)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d21_lattice3])])])])])])).
fof(c_0_10, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r2_hidden(X2,X3)=>(r3_lattices(X1,X2,k15_lattice3(X1,X3))&r3_lattices(X1,k16_lattice3(X1,X3),X2)))))), inference(assume_negation,[status(cth)],[t38_lattice3])).
cnf(c_0_11, plain, (r1_lattices(X2,X4,X1)|v2_struct_0(X2)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r4_lattice3(X2,X1,X3)|X4!=k15_lattice3(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_12, plain, ![X14, X15]:(v2_struct_0(X14)|~l3_lattices(X14)|m1_subset_1(k15_lattice3(X14,X15),u1_struct_0(X14))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k15_lattice3])])])).
fof(c_0_13, negated_conjecture, ((~v2_struct_0(esk1_0)&(v10_lattices(esk1_0)&(v4_lattice3(esk1_0)&l3_lattices(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(r2_hidden(esk2_0,esk3_0)&(~r3_lattices(esk1_0,esk2_0,k15_lattice3(esk1_0,esk3_0))|~r3_lattices(esk1_0,k16_lattice3(esk1_0,esk3_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_10])])])])).
fof(c_0_14, plain, ![X35, X36, X37]:((~r3_lattices(X35,X36,X37)|r1_lattices(X35,X36,X37)|(v2_struct_0(X35)|(~v6_lattices(X35)|(~v8_lattices(X35)|(~v9_lattices(X35)|~l3_lattices(X35))))|(~m1_subset_1(X36,u1_struct_0(X35))|~m1_subset_1(X37,u1_struct_0(X35)))))&(~r1_lattices(X35,X36,X37)|r3_lattices(X35,X36,X37)|(v2_struct_0(X35)|(~v6_lattices(X35)|(~v8_lattices(X35)|(~v9_lattices(X35)|~l3_lattices(X35))))|(~m1_subset_1(X36,u1_struct_0(X35))|~m1_subset_1(X37,u1_struct_0(X35)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_lattices])])])])).
cnf(c_0_15, plain, (v2_struct_0(X2)|r1_lattices(X2,X4,X1)|X4!=k15_lattice3(X2,X3)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)|~r4_lattice3(X2,X1,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))), inference(cn,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (v2_struct_0(X1)|m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_17, plain, ![X27, X28]:(v2_struct_0(X27)|~l3_lattices(X27)|k16_lattice3(X27,X28)=k15_lattice3(X27,a_2_1_lattice3(X27,X28))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d22_lattice3])])])])).
cnf(c_0_18, negated_conjecture, (~r3_lattices(esk1_0,esk2_0,k15_lattice3(esk1_0,esk3_0))|~r3_lattices(esk1_0,k16_lattice3(esk1_0,esk3_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (r3_lattices(X1,X2,X3)|v2_struct_0(X1)|~r1_lattices(X1,X2,X3)|~v6_lattices(X1)|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, negated_conjecture, (l3_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23, plain, (r1_lattices(X1,k15_lattice3(X1,X2),X3)|v2_struct_0(X1)|~r4_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_15]), c_0_16])).
cnf(c_0_24, plain, (v2_struct_0(X1)|k16_lattice3(X1,X2)=k15_lattice3(X1,a_2_1_lattice3(X1,X2))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_25, plain, ![X29, X30, X31, X33, X34]:(((m1_subset_1(esk7_3(X29,X30,X31),u1_struct_0(X30))|~r2_hidden(X29,a_2_1_lattice3(X30,X31))|(v2_struct_0(X30)|~l3_lattices(X30)))&((X29=esk7_3(X29,X30,X31)|~r2_hidden(X29,a_2_1_lattice3(X30,X31))|(v2_struct_0(X30)|~l3_lattices(X30)))&(r3_lattice3(X30,esk7_3(X29,X30,X31),X31)|~r2_hidden(X29,a_2_1_lattice3(X30,X31))|(v2_struct_0(X30)|~l3_lattices(X30)))))&(~m1_subset_1(X34,u1_struct_0(X30))|(X29!=X34|~r3_lattice3(X30,X34,X33))|r2_hidden(X29,a_2_1_lattice3(X30,X33))|(v2_struct_0(X30)|~l3_lattices(X30)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_1_lattice3])])])])])])])).
fof(c_0_26, plain, ![X8, X9, X10, X11, X12]:((~r4_lattice3(X8,X9,X10)|(~m1_subset_1(X11,u1_struct_0(X8))|(~r2_hidden(X11,X10)|r1_lattices(X8,X11,X9)))|~m1_subset_1(X9,u1_struct_0(X8))|(v2_struct_0(X8)|~l3_lattices(X8)))&((m1_subset_1(esk4_3(X8,X9,X12),u1_struct_0(X8))|r4_lattice3(X8,X9,X12)|~m1_subset_1(X9,u1_struct_0(X8))|(v2_struct_0(X8)|~l3_lattices(X8)))&((r2_hidden(esk4_3(X8,X9,X12),X12)|r4_lattice3(X8,X9,X12)|~m1_subset_1(X9,u1_struct_0(X8))|(v2_struct_0(X8)|~l3_lattices(X8)))&(~r1_lattices(X8,esk4_3(X8,X9,X12),X9)|r4_lattice3(X8,X9,X12)|~m1_subset_1(X9,u1_struct_0(X8))|(v2_struct_0(X8)|~l3_lattices(X8)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_lattice3])])])])])])])).
cnf(c_0_27, negated_conjecture, (~v9_lattices(esk1_0)|~v8_lattices(esk1_0)|~v6_lattices(esk1_0)|~r1_lattices(esk1_0,k16_lattice3(esk1_0,esk3_0),esk2_0)|~r3_lattices(esk1_0,esk2_0,k15_lattice3(esk1_0,esk3_0))|~m1_subset_1(k16_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_28, plain, (r1_lattices(X1,k16_lattice3(X1,X2),X3)|v2_struct_0(X1)|~r4_lattice3(X1,X3,a_2_1_lattice3(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_29, negated_conjecture, (v4_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_30, negated_conjecture, (v10_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_31, plain, (r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X1)|X2!=k15_lattice3(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_32, plain, (X1=esk7_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_lattice3(X2,X3))|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (r2_hidden(esk4_3(X1,X2,X3),X3)|r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, negated_conjecture, (~v9_lattices(esk1_0)|~v8_lattices(esk1_0)|~v6_lattices(esk1_0)|~r4_lattice3(esk1_0,esk2_0,a_2_1_lattice3(esk1_0,esk3_0))|~r3_lattices(esk1_0,esk2_0,k15_lattice3(esk1_0,esk3_0))|~m1_subset_1(k16_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_20]), c_0_21]), c_0_29]), c_0_30])]), c_0_22])).
cnf(c_0_35, plain, (v2_struct_0(X1)|r4_lattice3(X1,X2,X3)|X2!=k15_lattice3(X1,X3)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_31])).
fof(c_0_36, plain, ![X21, X22, X23, X24, X25]:((~r3_lattice3(X21,X22,X23)|(~m1_subset_1(X24,u1_struct_0(X21))|(~r2_hidden(X24,X23)|r1_lattices(X21,X22,X24)))|~m1_subset_1(X22,u1_struct_0(X21))|(v2_struct_0(X21)|~l3_lattices(X21)))&((m1_subset_1(esk6_3(X21,X22,X25),u1_struct_0(X21))|r3_lattice3(X21,X22,X25)|~m1_subset_1(X22,u1_struct_0(X21))|(v2_struct_0(X21)|~l3_lattices(X21)))&((r2_hidden(esk6_3(X21,X22,X25),X25)|r3_lattice3(X21,X22,X25)|~m1_subset_1(X22,u1_struct_0(X21))|(v2_struct_0(X21)|~l3_lattices(X21)))&(~r1_lattices(X21,X22,esk6_3(X21,X22,X25))|r3_lattice3(X21,X22,X25)|~m1_subset_1(X22,u1_struct_0(X21))|(v2_struct_0(X21)|~l3_lattices(X21)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattice3])])])])])])])).
cnf(c_0_37, plain, (r3_lattice3(X1,esk7_3(X2,X1,X3),X3)|v2_struct_0(X1)|~r2_hidden(X2,a_2_1_lattice3(X1,X3))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_38, plain, (esk7_3(esk4_3(X1,X2,a_2_1_lattice3(X3,X4)),X3,X4)=esk4_3(X1,X2,a_2_1_lattice3(X3,X4))|r4_lattice3(X1,X2,a_2_1_lattice3(X3,X4))|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X3)|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_39, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_lattice3(X2,X3))|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_40, negated_conjecture, (~v9_lattices(esk1_0)|~v8_lattices(esk1_0)|~v6_lattices(esk1_0)|~r1_lattices(esk1_0,esk2_0,k15_lattice3(esk1_0,esk3_0))|~r4_lattice3(esk1_0,esk2_0,a_2_1_lattice3(esk1_0,esk3_0))|~m1_subset_1(k16_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k15_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_19]), c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_41, plain, (m1_subset_1(k16_lattice3(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_16, c_0_24])).
cnf(c_0_42, plain, (r1_lattices(X1,X4,X2)|v2_struct_0(X1)|~r4_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_43, plain, (r4_lattice3(X1,k15_lattice3(X1,X2),X2)|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_35]), c_0_16])).
cnf(c_0_44, plain, (r1_lattices(X1,X2,X4)|v2_struct_0(X1)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_45, plain, (r3_lattice3(X1,esk4_3(X2,X3,a_2_1_lattice3(X1,X4)),X4)|r4_lattice3(X2,X3,a_2_1_lattice3(X1,X4))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))|~l3_lattices(X1)|~l3_lattices(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_33])).
cnf(c_0_46, plain, (r4_lattice3(X1,X2,a_2_1_lattice3(X3,X4))|m1_subset_1(esk4_3(X1,X2,a_2_1_lattice3(X3,X4)),u1_struct_0(X3))|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X3)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_38]), c_0_33])).
cnf(c_0_47, negated_conjecture, (~v9_lattices(esk1_0)|~v8_lattices(esk1_0)|~v6_lattices(esk1_0)|~r1_lattices(esk1_0,esk2_0,k15_lattice3(esk1_0,esk3_0))|~r4_lattice3(esk1_0,esk2_0,a_2_1_lattice3(esk1_0,esk3_0))|~m1_subset_1(k15_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_21])]), c_0_22])).
cnf(c_0_48, plain, (r1_lattices(X1,X2,k15_lattice3(X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_16])).
cnf(c_0_49, negated_conjecture, (r2_hidden(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_50, plain, (r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r1_lattices(X1,esk4_3(X1,X2,X3),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_51, plain, (r1_lattices(X1,esk4_3(X2,X3,a_2_1_lattice3(X1,X4)),X5)|r4_lattice3(X2,X3,a_2_1_lattice3(X1,X4))|v2_struct_0(X2)|v2_struct_0(X1)|~r2_hidden(X5,X4)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_lattices(X1)|~l3_lattices(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])).
cnf(c_0_52, negated_conjecture, (~v9_lattices(esk1_0)|~v8_lattices(esk1_0)|~v6_lattices(esk1_0)|~r4_lattice3(esk1_0,esk2_0,a_2_1_lattice3(esk1_0,esk3_0))|~m1_subset_1(k15_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49]), c_0_20]), c_0_21]), c_0_29]), c_0_30])]), c_0_22])).
cnf(c_0_53, plain, (r4_lattice3(X1,X2,a_2_1_lattice3(X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_54, negated_conjecture, (~v9_lattices(esk1_0)|~v8_lattices(esk1_0)|~v6_lattices(esk1_0)|~m1_subset_1(k15_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_49]), c_0_20]), c_0_21])]), c_0_22])).
fof(c_0_55, plain, ![X38]:((~v2_struct_0(X38)|(v2_struct_0(X38)|~v10_lattices(X38))|~l3_lattices(X38))&((v4_lattices(X38)|(v2_struct_0(X38)|~v10_lattices(X38))|~l3_lattices(X38))&((v5_lattices(X38)|(v2_struct_0(X38)|~v10_lattices(X38))|~l3_lattices(X38))&((v6_lattices(X38)|(v2_struct_0(X38)|~v10_lattices(X38))|~l3_lattices(X38))&((v7_lattices(X38)|(v2_struct_0(X38)|~v10_lattices(X38))|~l3_lattices(X38))&((v8_lattices(X38)|(v2_struct_0(X38)|~v10_lattices(X38))|~l3_lattices(X38))&(v9_lattices(X38)|(v2_struct_0(X38)|~v10_lattices(X38))|~l3_lattices(X38)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
cnf(c_0_56, negated_conjecture, (~v9_lattices(esk1_0)|~v8_lattices(esk1_0)|~v6_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_16]), c_0_21])]), c_0_22])).
cnf(c_0_57, plain, (v9_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_58, negated_conjecture, (~v8_lattices(esk1_0)|~v6_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_21]), c_0_30])]), c_0_22])).
cnf(c_0_59, plain, (v8_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_60, negated_conjecture, (~v6_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_21]), c_0_30])]), c_0_22])).
cnf(c_0_61, plain, (v6_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_62, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_21]), c_0_30])]), c_0_22]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 63
# Proof object clause steps            : 44
# Proof object formula steps           : 19
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 54
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 34
# Processed clauses                    : 129
# ...of these trivial                  : 0
# ...subsumed                          : 15
# ...remaining for further processing  : 114
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 9
# Backward-rewritten                   : 0
# Generated clauses                    : 94
# ...of the previous two non-trivial   : 88
# Contextual simplify-reflections      : 21
# Paramodulations                      : 91
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 68
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 61
# Current number of unprocessed clauses: 26
# ...number of literals in the above   : 257
# Current number of archived formulas  : 0
# Current number of archived clauses   : 43
# Clause-clause subsumption calls (NU) : 3480
# Rec. Clause-clause subsumption calls : 114
# Non-unit clause-clause subsumptions  : 45
# Unit Clause-clause subsumption calls : 62
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7016

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.000 s
# Total time               : 0.015 s
# Maximum resident set size: 28684 pages
