# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_ordinal1, axiom, ![X1]:(v1_ordinal1(X1)<=>![X2]:(r2_hidden(X2,X1)=>r1_tarski(X2,X1))), file('data/MPTP/proofs/5b537e3e762945098daaf91f883d4f17.E_input', d2_ordinal1)).
fof(cc1_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>(v1_ordinal1(X1)&v2_ordinal1(X1))), file('data/MPTP/proofs/5b537e3e762945098daaf91f883d4f17.E_input', cc1_ordinal1)).
fof(t33_ordinal1, conjecture, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)<=>r1_ordinal1(k1_ordinal1(X1),X2)))), file('data/MPTP/proofs/5b537e3e762945098daaf91f883d4f17.E_input', t33_ordinal1)).
fof(t13_ordinal1, axiom, ![X1, X2]:(r2_hidden(X1,k1_ordinal1(X2))<=>(r2_hidden(X1,X2)|X1=X2)), file('data/MPTP/proofs/5b537e3e762945098daaf91f883d4f17.E_input', t13_ordinal1)).
fof(d1_ordinal1, axiom, ![X1]:k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1)), file('data/MPTP/proofs/5b537e3e762945098daaf91f883d4f17.E_input', d1_ordinal1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/5b537e3e762945098daaf91f883d4f17.E_input', d3_tarski)).
fof(redefinition_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)<=>r1_tarski(X1,X2))), file('data/MPTP/proofs/5b537e3e762945098daaf91f883d4f17.E_input', redefinition_r1_ordinal1)).
fof(t29_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>v3_ordinal1(k1_ordinal1(X1))), file('data/MPTP/proofs/5b537e3e762945098daaf91f883d4f17.E_input', t29_ordinal1)).
fof(t10_ordinal1, axiom, ![X1]:r2_hidden(X1,k1_ordinal1(X1)), file('data/MPTP/proofs/5b537e3e762945098daaf91f883d4f17.E_input', t10_ordinal1)).
fof(c_0_9, plain, ![X11, X12, X13]:((~v1_ordinal1(X11)|(~r2_hidden(X12,X11)|r1_tarski(X12,X11)))&((r2_hidden(esk3_1(X13),X13)|v1_ordinal1(X13))&(~r1_tarski(esk3_1(X13),X13)|v1_ordinal1(X13)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_ordinal1])])])])])])).
fof(c_0_10, plain, ![X10]:((v1_ordinal1(X10)|~v3_ordinal1(X10))&(v2_ordinal1(X10)|~v3_ordinal1(X10))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_ordinal1])])])).
fof(c_0_11, negated_conjecture, ~(![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)<=>r1_ordinal1(k1_ordinal1(X1),X2))))), inference(assume_negation,[status(cth)],[t33_ordinal1])).
fof(c_0_12, plain, ![X6, X7]:((~r2_hidden(X6,k1_ordinal1(X7))|(r2_hidden(X6,X7)|X6=X7))&((~r2_hidden(X6,X7)|r2_hidden(X6,k1_ordinal1(X7)))&(X6!=X7|r2_hidden(X6,k1_ordinal1(X7))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_ordinal1])])])).
fof(c_0_13, plain, ![X15]:k1_ordinal1(X15)=k2_xboole_0(X15,k1_tarski(X15)), inference(variable_rename,[status(thm)],[d1_ordinal1])).
fof(c_0_14, plain, ![X16, X17, X18, X19, X20]:((~r1_tarski(X16,X17)|(~r2_hidden(X18,X16)|r2_hidden(X18,X17)))&((r2_hidden(esk4_2(X19,X20),X19)|r1_tarski(X19,X20))&(~r2_hidden(esk4_2(X19,X20),X20)|r1_tarski(X19,X20)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_15, plain, (r1_tarski(X2,X1)|~v1_ordinal1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, plain, (v1_ordinal1(X1)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_17, negated_conjecture, (v3_ordinal1(esk1_0)&(v3_ordinal1(esk2_0)&((~r2_hidden(esk1_0,esk2_0)|~r1_ordinal1(k1_ordinal1(esk1_0),esk2_0))&(r2_hidden(esk1_0,esk2_0)|r1_ordinal1(k1_ordinal1(esk1_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
cnf(c_0_18, plain, (r2_hidden(X1,k1_ordinal1(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,X2)|~v3_ordinal1(X2)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
fof(c_0_22, plain, ![X22, X23]:((~r1_ordinal1(X22,X23)|r1_tarski(X22,X23)|(~v3_ordinal1(X22)|~v3_ordinal1(X23)))&(~r1_tarski(X22,X23)|r1_ordinal1(X22,X23)|(~v3_ordinal1(X22)|~v3_ordinal1(X23)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r1_ordinal1])])])).
cnf(c_0_23, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|r1_ordinal1(k1_ordinal1(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (r1_tarski(X1,X2)|~r2_hidden(esk4_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_25, plain, (r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_26, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r2_hidden(X3,X2)|~v3_ordinal1(X2)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_28, plain, (r1_tarski(X1,X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|r1_ordinal1(k2_xboole_0(esk1_0,k1_tarski(esk1_0)),esk2_0)), inference(rw,[status(thm)],[c_0_23, c_0_19])).
cnf(c_0_30, negated_conjecture, (v3_ordinal1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_31, plain, ![X8]:(~v3_ordinal1(X8)|v3_ordinal1(k1_ordinal1(X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t29_ordinal1])])).
cnf(c_0_32, plain, (r1_tarski(X1,k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(esk4_2(X1,k2_xboole_0(X2,k1_tarski(X2))),X2)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_33, plain, (r1_tarski(X1,X2)|r2_hidden(esk4_2(X1,X2),X3)|~r2_hidden(X1,X3)|~v3_ordinal1(X3)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_34, plain, (r2_hidden(X1,X2)|X1=X2|~r2_hidden(X1,k1_ordinal1(X2))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_35, negated_conjecture, (r1_tarski(k2_xboole_0(esk1_0,k1_tarski(esk1_0)),esk2_0)|r2_hidden(esk1_0,esk2_0)|~v3_ordinal1(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])])).
cnf(c_0_36, plain, (v3_ordinal1(k1_ordinal1(X1))|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_37, plain, ![X9]:r2_hidden(X9,k1_ordinal1(X9)), inference(variable_rename,[status(thm)],[t10_ordinal1])).
cnf(c_0_38, plain, (v1_ordinal1(X1)|~r1_tarski(esk3_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_39, plain, (r1_tarski(X1,k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(X1,X2)|~v3_ordinal1(X2)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_40, plain, (X1=X2|r2_hidden(X1,X2)|~r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))), inference(rw,[status(thm)],[c_0_34, c_0_19])).
cnf(c_0_41, plain, (r2_hidden(esk3_1(X1),X1)|v1_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_42, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|r2_hidden(X1,esk2_0)|~r2_hidden(X1,k2_xboole_0(esk1_0,k1_tarski(esk1_0)))|~v3_ordinal1(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))), inference(spm,[status(thm)],[c_0_20, c_0_35])).
cnf(c_0_43, plain, (v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)), inference(rw,[status(thm)],[c_0_36, c_0_19])).
cnf(c_0_44, negated_conjecture, (v3_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_45, plain, (r2_hidden(X1,k1_ordinal1(X1))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_46, plain, (v1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(esk3_1(k2_xboole_0(X1,k1_tarski(X1))),X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_47, plain, (esk3_1(k2_xboole_0(X1,k1_tarski(X1)))=X1|v1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|r2_hidden(esk3_1(k2_xboole_0(X1,k1_tarski(X1))),X1)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_48, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|r2_hidden(X1,esk2_0)|~r2_hidden(X1,k2_xboole_0(esk1_0,k1_tarski(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])])).
cnf(c_0_49, plain, (r2_hidden(X1,k2_xboole_0(X1,k1_tarski(X1)))), inference(rw,[status(thm)],[c_0_45, c_0_19])).
cnf(c_0_50, plain, (esk3_1(k2_xboole_0(X1,k1_tarski(X1)))=X1|v1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_51, plain, (r1_tarski(X1,k2_xboole_0(X1,k1_tarski(X1)))), inference(spm,[status(thm)],[c_0_32, c_0_27])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_53, plain, (v1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_50]), c_0_51])])).
cnf(c_0_54, negated_conjecture, (r2_hidden(esk1_0,X1)|~r2_hidden(esk2_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_26, c_0_52])).
cnf(c_0_55, plain, (r1_tarski(X1,k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X2)), inference(spm,[status(thm)],[c_0_15, c_0_53])).
cnf(c_0_56, negated_conjecture, (r2_hidden(esk1_0,k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(esk2_0,k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_54, c_0_43])).
cnf(c_0_57, plain, (r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(X3,k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(X1,X3)|~v3_ordinal1(X2)), inference(spm,[status(thm)],[c_0_20, c_0_55])).
cnf(c_0_58, negated_conjecture, (r2_hidden(esk1_0,k2_xboole_0(esk2_0,k1_tarski(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_49]), c_0_30])])).
cnf(c_0_59, negated_conjecture, (r2_hidden(X1,k2_xboole_0(esk2_0,k1_tarski(esk2_0)))|~r2_hidden(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_30])])).
cnf(c_0_60, negated_conjecture, (X1=esk2_0|r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_40, c_0_59])).
cnf(c_0_61, negated_conjecture, (esk4_2(X1,esk2_0)=esk2_0|r1_tarski(X1,esk2_0)|~r2_hidden(esk4_2(X1,esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_24, c_0_60])).
cnf(c_0_62, negated_conjecture, (esk4_2(esk1_0,esk2_0)=esk2_0|r1_tarski(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_61, c_0_27])).
cnf(c_0_63, negated_conjecture, (r1_tarski(esk1_0,esk2_0)|r2_hidden(esk2_0,X1)|~r2_hidden(esk1_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_33, c_0_62])).
cnf(c_0_64, negated_conjecture, (r1_tarski(esk1_0,esk2_0)|~r2_hidden(esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_24, c_0_62])).
cnf(c_0_65, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_30]), c_0_52])]), c_0_64])).
cnf(c_0_66, negated_conjecture, (~r2_hidden(esk1_0,esk2_0)|~r1_ordinal1(k1_ordinal1(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_67, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_20, c_0_65])).
cnf(c_0_68, negated_conjecture, (~r2_hidden(esk1_0,esk2_0)|~r1_ordinal1(k2_xboole_0(esk1_0,k1_tarski(esk1_0)),esk2_0)), inference(rw,[status(thm)],[c_0_66, c_0_19])).
cnf(c_0_69, plain, (r1_ordinal1(X1,X2)|~r1_tarski(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_70, negated_conjecture, (r1_tarski(X1,esk2_0)|~r2_hidden(esk4_2(X1,esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_24, c_0_67])).
cnf(c_0_71, plain, (esk4_2(k2_xboole_0(X1,k1_tarski(X1)),X2)=X1|r1_tarski(k2_xboole_0(X1,k1_tarski(X1)),X2)|r2_hidden(esk4_2(k2_xboole_0(X1,k1_tarski(X1)),X2),X1)), inference(spm,[status(thm)],[c_0_40, c_0_27])).
cnf(c_0_72, negated_conjecture, (~r1_tarski(k2_xboole_0(esk1_0,k1_tarski(esk1_0)),esk2_0)|~r2_hidden(esk1_0,esk2_0)|~v3_ordinal1(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_30])])).
cnf(c_0_73, negated_conjecture, (esk4_2(k2_xboole_0(esk1_0,k1_tarski(esk1_0)),esk2_0)=esk1_0|r1_tarski(k2_xboole_0(esk1_0,k1_tarski(esk1_0)),esk2_0)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_74, negated_conjecture, (~r1_tarski(k2_xboole_0(esk1_0,k1_tarski(esk1_0)),esk2_0)|~v3_ordinal1(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_52])])).
cnf(c_0_75, negated_conjecture, (r1_tarski(k2_xboole_0(esk1_0,k1_tarski(esk1_0)),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_73]), c_0_52])])).
cnf(c_0_76, negated_conjecture, (~v3_ordinal1(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_75])])).
cnf(c_0_77, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_43]), c_0_44])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 78
# Proof object clause steps            : 59
# Proof object formula steps           : 19
# Proof object conjectures             : 31
# Proof object clause conjectures      : 28
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 33
# Proof object simplifying inferences  : 29
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 19
# Processed clauses                    : 252
# ...of these trivial                  : 0
# ...subsumed                          : 62
# ...remaining for further processing  : 190
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 23
# Backward-rewritten                   : 10
# Generated clauses                    : 807
# ...of the previous two non-trivial   : 753
# Contextual simplify-reflections      : 5
# Paramodulations                      : 806
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 138
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 125
# Current number of unprocessed clauses: 529
# ...number of literals in the above   : 2884
# Current number of archived formulas  : 0
# Current number of archived clauses   : 52
# Clause-clause subsumption calls (NU) : 3180
# Rec. Clause-clause subsumption calls : 1800
# Non-unit clause-clause subsumptions  : 83
# Unit Clause-clause subsumption calls : 24
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 18898

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.003 s
# Total time               : 0.020 s
# Maximum resident set size: 28736 pages
