# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_waybel_7, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v1_subset_1(X2,u1_struct_0(X1))<=>~(r2_hidden(k3_yellow_0(X1),X2))))), file('data/MPTP/proofs/5c267ee5adeb4f0eb3a8fa1a6d4ca5ce.E_input', t8_waybel_7)).
fof(t34_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:((r1_tarski(X2,X3)&(r1_yellow_0(X1,X2)&r1_yellow_0(X1,X3)))=>r1_orders_2(X1,k1_yellow_0(X1,X2),k1_yellow_0(X1,X3)))), file('data/MPTP/proofs/5c267ee5adeb4f0eb3a8fa1a6d4ca5ce.E_input', t34_yellow_0)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/5c267ee5adeb4f0eb3a8fa1a6d4ca5ce.E_input', t2_xboole_1)).
fof(t42_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))=>(r1_yellow_0(X1,k1_xboole_0)&r2_yellow_0(X1,u1_struct_0(X1)))), file('data/MPTP/proofs/5c267ee5adeb4f0eb3a8fa1a6d4ca5ce.E_input', t42_yellow_0)).
fof(t49_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(![X3]:(m1_subset_1(X3,X1)=>r2_hidden(X3,X2))=>X1=X2)), file('data/MPTP/proofs/5c267ee5adeb4f0eb3a8fa1a6d4ca5ce.E_input', t49_subset_1)).
fof(t34_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r1_yellow_0(X1,k5_waybel_0(X1,X2))&k1_yellow_0(X1,k5_waybel_0(X1,X2))=X2))), file('data/MPTP/proofs/5c267ee5adeb4f0eb3a8fa1a6d4ca5ce.E_input', t34_waybel_0)).
fof(d20_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v13_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r1_orders_2(X1,X3,X4))=>r2_hidden(X4,X2))))))), file('data/MPTP/proofs/5c267ee5adeb4f0eb3a8fa1a6d4ca5ce.E_input', d20_waybel_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/5c267ee5adeb4f0eb3a8fa1a6d4ca5ce.E_input', t4_subset)).
fof(d7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(v1_subset_1(X2,X1)<=>~(X2=X1))), file('data/MPTP/proofs/5c267ee5adeb4f0eb3a8fa1a6d4ca5ce.E_input', d7_subset_1)).
fof(d11_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0)), file('data/MPTP/proofs/5c267ee5adeb4f0eb3a8fa1a6d4ca5ce.E_input', d11_yellow_0)).
fof(fc14_subset_1, axiom, ![X1]:~(v1_subset_1(k2_subset_1(X1),X1)), file('data/MPTP/proofs/5c267ee5adeb4f0eb3a8fa1a6d4ca5ce.E_input', fc14_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/5c267ee5adeb4f0eb3a8fa1a6d4ca5ce.E_input', d4_subset_1)).
fof(dt_k3_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(k3_yellow_0(X1),u1_struct_0(X1))), file('data/MPTP/proofs/5c267ee5adeb4f0eb3a8fa1a6d4ca5ce.E_input', dt_k3_yellow_0)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/5c267ee5adeb4f0eb3a8fa1a6d4ca5ce.E_input', t2_subset)).
fof(c_0_14, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v1_subset_1(X2,u1_struct_0(X1))<=>~(r2_hidden(k3_yellow_0(X1),X2)))))), inference(assume_negation,[status(cth)],[t8_waybel_7])).
fof(c_0_15, plain, ![X21, X22, X23]:(~l1_orders_2(X21)|(~r1_tarski(X22,X23)|(~r1_yellow_0(X21,X22)|~r1_yellow_0(X21,X23))|r1_orders_2(X21,k1_yellow_0(X21,X22),k1_yellow_0(X21,X23)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t34_yellow_0])])])).
fof(c_0_16, plain, ![X25]:r1_tarski(k1_xboole_0,X25), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_17, plain, ![X7]:((r1_yellow_0(X7,k1_xboole_0)|(v2_struct_0(X7)|(~v5_orders_2(X7)|(~v1_yellow_0(X7)|~l1_orders_2(X7)))))&(r2_yellow_0(X7,u1_struct_0(X7))|(v2_struct_0(X7)|(~v5_orders_2(X7)|(~v1_yellow_0(X7)|~l1_orders_2(X7)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t42_yellow_0])])])])).
fof(c_0_18, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_yellow_0(esk1_0)&l1_orders_2(esk1_0))))))&((~v1_xboole_0(esk2_0)&(v2_waybel_0(esk2_0,esk1_0)&(v13_waybel_0(esk2_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))))&((~v1_subset_1(esk2_0,u1_struct_0(esk1_0))|r2_hidden(k3_yellow_0(esk1_0),esk2_0))&(v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k3_yellow_0(esk1_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])).
fof(c_0_19, plain, ![X28, X29]:((m1_subset_1(esk5_2(X28,X29),X28)|X28=X29|~m1_subset_1(X29,k1_zfmisc_1(X28)))&(~r2_hidden(esk5_2(X28,X29),X29)|X28=X29|~m1_subset_1(X29,k1_zfmisc_1(X28)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t49_subset_1])])])])).
cnf(c_0_20, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),k1_yellow_0(X1,X3))|~l1_orders_2(X1)|~r1_tarski(X2,X3)|~r1_yellow_0(X1,X2)|~r1_yellow_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (r1_yellow_0(X1,k1_xboole_0)|v2_struct_0(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, negated_conjecture, (v1_yellow_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_27, plain, ![X32, X33]:((r1_yellow_0(X32,k5_waybel_0(X32,X33))|~m1_subset_1(X33,u1_struct_0(X32))|(v2_struct_0(X32)|(~v3_orders_2(X32)|(~v4_orders_2(X32)|(~v5_orders_2(X32)|~l1_orders_2(X32))))))&(k1_yellow_0(X32,k5_waybel_0(X32,X33))=X33|~m1_subset_1(X33,u1_struct_0(X32))|(v2_struct_0(X32)|(~v3_orders_2(X32)|(~v4_orders_2(X32)|(~v5_orders_2(X32)|~l1_orders_2(X32))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_waybel_0])])])])])).
cnf(c_0_28, plain, (m1_subset_1(esk5_2(X1,X2),X1)|X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_30, plain, ![X12, X13, X14, X15]:((~v13_waybel_0(X13,X12)|(~m1_subset_1(X14,u1_struct_0(X12))|(~m1_subset_1(X15,u1_struct_0(X12))|(~r2_hidden(X14,X13)|~r1_orders_2(X12,X14,X15)|r2_hidden(X15,X13))))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12))&((m1_subset_1(esk3_2(X12,X13),u1_struct_0(X12))|v13_waybel_0(X13,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12))&((m1_subset_1(esk4_2(X12,X13),u1_struct_0(X12))|v13_waybel_0(X13,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12))&(((r2_hidden(esk3_2(X12,X13),X13)|v13_waybel_0(X13,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12))&(r1_orders_2(X12,esk3_2(X12,X13),esk4_2(X12,X13))|v13_waybel_0(X13,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12)))&(~r2_hidden(esk4_2(X12,X13),X13)|v13_waybel_0(X13,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d20_waybel_0])])])])])).
fof(c_0_31, plain, ![X9, X10, X11]:(~r2_hidden(X9,X10)|~m1_subset_1(X10,k1_zfmisc_1(X11))|m1_subset_1(X9,X11)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_32, plain, (r1_orders_2(X1,k1_yellow_0(X1,k1_xboole_0),k1_yellow_0(X1,X2))|~r1_yellow_0(X1,k1_xboole_0)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_33, negated_conjecture, (r1_yellow_0(esk1_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_34, plain, (r1_yellow_0(X1,k5_waybel_0(X1,X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|m1_subset_1(esk5_2(u1_struct_0(esk1_0),esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_36, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_37, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_38, plain, (r2_hidden(X4,X1)|~v13_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,X1)|~r1_orders_2(X2,X3,X4)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_40, plain, ![X19, X20]:((~v1_subset_1(X20,X19)|X20!=X19|~m1_subset_1(X20,k1_zfmisc_1(X19)))&(X20=X19|v1_subset_1(X20,X19)|~m1_subset_1(X20,k1_zfmisc_1(X19)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_subset_1])])])])).
cnf(c_0_41, negated_conjecture, (r1_orders_2(esk1_0,k1_yellow_0(esk1_0,k1_xboole_0),k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_24])])).
cnf(c_0_42, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r1_yellow_0(esk1_0,k5_waybel_0(esk1_0,esk5_2(u1_struct_0(esk1_0),esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_24]), c_0_25]), c_0_36]), c_0_37])]), c_0_26])).
cnf(c_0_43, plain, (k1_yellow_0(X1,k5_waybel_0(X1,X2))=X2|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_44, plain, (r2_hidden(X1,X2)|~r1_orders_2(X3,X4,X1)|~r2_hidden(X4,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~v13_waybel_0(X2,X3)|~l1_orders_2(X3)), inference(csr,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_45, plain, (X1=X2|~r2_hidden(esk5_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_46, plain, (X1=X2|v1_subset_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_47, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,k1_xboole_0),k1_yellow_0(esk1_0,k5_waybel_0(esk1_0,esk5_2(u1_struct_0(esk1_0),esk2_0))))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_48, negated_conjecture, (k1_yellow_0(esk1_0,k5_waybel_0(esk1_0,esk5_2(u1_struct_0(esk1_0),esk2_0)))=esk5_2(u1_struct_0(esk1_0),esk2_0)|u1_struct_0(esk1_0)=esk2_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_35]), c_0_24]), c_0_25]), c_0_36]), c_0_37])]), c_0_26])).
fof(c_0_49, plain, ![X8]:(~l1_orders_2(X8)|k3_yellow_0(X8)=k1_yellow_0(X8,k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_yellow_0])])).
cnf(c_0_50, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_hidden(esk5_2(u1_struct_0(esk1_0),esk2_0),X1)|~r1_orders_2(esk1_0,X2,esk5_2(u1_struct_0(esk1_0),esk2_0))|~r2_hidden(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v13_waybel_0(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_35]), c_0_24])])).
cnf(c_0_51, negated_conjecture, (v13_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_52, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|~r2_hidden(esk5_2(u1_struct_0(esk1_0),esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_45, c_0_29])).
cnf(c_0_53, negated_conjecture, (r2_hidden(k3_yellow_0(esk1_0),esk2_0)|~v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_54, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_46, c_0_29])).
cnf(c_0_55, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,k1_xboole_0),esk5_2(u1_struct_0(esk1_0),esk2_0))), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_56, plain, (k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_57, plain, ![X24]:~v1_subset_1(k2_subset_1(X24),X24), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc14_subset_1])])).
fof(c_0_58, plain, ![X18]:k2_subset_1(X18)=X18, inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_59, plain, ![X31]:(~l1_orders_2(X31)|m1_subset_1(k3_yellow_0(X31),u1_struct_0(X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_yellow_0])])).
cnf(c_0_60, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|~r1_orders_2(esk1_0,X1,esk5_2(u1_struct_0(esk1_0),esk2_0))|~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_29])]), c_0_52])).
cnf(c_0_61, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_hidden(k3_yellow_0(esk1_0),esk2_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_62, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r1_orders_2(esk1_0,k3_yellow_0(esk1_0),esk5_2(u1_struct_0(esk1_0),esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_24])])).
cnf(c_0_63, plain, (~v1_subset_1(k2_subset_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_64, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_58])).
fof(c_0_65, plain, ![X26, X27]:(~m1_subset_1(X26,X27)|(v1_xboole_0(X27)|r2_hidden(X26,X27))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_66, plain, (m1_subset_1(k3_yellow_0(X1),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_67, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62])).
cnf(c_0_68, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k3_yellow_0(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_69, plain, (~v1_subset_1(X1,X1)), inference(rw,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_70, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_71, negated_conjecture, (m1_subset_1(k3_yellow_0(esk1_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_24])])).
cnf(c_0_72, negated_conjecture, (~r2_hidden(k3_yellow_0(esk1_0),esk2_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_67]), c_0_69])).
cnf(c_0_73, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_74, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_72]), c_0_73]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 75
# Proof object clause steps            : 46
# Proof object formula steps           : 29
# Proof object conjectures             : 31
# Proof object clause conjectures      : 28
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 34
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 33
# Processed clauses                    : 110
# ...of these trivial                  : 2
# ...subsumed                          : 5
# ...remaining for further processing  : 103
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 23
# Generated clauses                    : 61
# ...of the previous two non-trivial   : 57
# Contextual simplify-reflections      : 3
# Paramodulations                      : 60
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 46
#    Positive orientable unit clauses  : 18
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 24
# Current number of unprocessed clauses: 12
# ...number of literals in the above   : 35
# Current number of archived formulas  : 0
# Current number of archived clauses   : 57
# Clause-clause subsumption calls (NU) : 409
# Rec. Clause-clause subsumption calls : 110
# Non-unit clause-clause subsumptions  : 7
# Unit Clause-clause subsumption calls : 35
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3705

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28732 pages
