# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t34_funct_2, conjecture, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X2,X1)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))))=>((k2_relset_1(X1,X2,X3)=X2&(v2_funct_1(X3)&![X5, X6]:(((r2_hidden(X5,X2)&k1_funct_1(X4,X5)=X6)=>(r2_hidden(X6,X1)&k1_funct_1(X3,X6)=X5))&((r2_hidden(X6,X1)&k1_funct_1(X3,X6)=X5)=>(r2_hidden(X5,X2)&k1_funct_1(X4,X5)=X6)))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|X4=k2_funct_1(X3)))))), file('data/MPTP/proofs/5c966c6bec174891b3008e1f15a158f0.E_input', t34_funct_2)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(((X2=k1_xboole_0=>X1=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X2,X3)))&(X2=k1_xboole_0=>(X1=k1_xboole_0|(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0))))), file('data/MPTP/proofs/5c966c6bec174891b3008e1f15a158f0.E_input', d1_funct_2)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/5c966c6bec174891b3008e1f15a158f0.E_input', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/5c966c6bec174891b3008e1f15a158f0.E_input', fc6_relat_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/5c966c6bec174891b3008e1f15a158f0.E_input', redefinition_k1_relset_1)).
fof(t54_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(X2=k2_funct_1(X1)<=>(k1_relat_1(X2)=k2_relat_1(X1)&![X3, X4]:(((r2_hidden(X3,k2_relat_1(X1))&X4=k1_funct_1(X2,X3))=>(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4)))&((r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4))=>(r2_hidden(X3,k2_relat_1(X1))&X4=k1_funct_1(X2,X3))))))))), file('data/MPTP/proofs/5c966c6bec174891b3008e1f15a158f0.E_input', t54_funct_1)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/5c966c6bec174891b3008e1f15a158f0.E_input', redefinition_k2_relset_1)).
fof(c_0_7, negated_conjecture, ~(![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X2,X1)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))))=>((k2_relset_1(X1,X2,X3)=X2&(v2_funct_1(X3)&![X5, X6]:(((r2_hidden(X5,X2)&k1_funct_1(X4,X5)=X6)=>(r2_hidden(X6,X1)&k1_funct_1(X3,X6)=X5))&((r2_hidden(X6,X1)&k1_funct_1(X3,X6)=X5)=>(r2_hidden(X5,X2)&k1_funct_1(X4,X5)=X6)))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|X4=k2_funct_1(X3))))))), inference(assume_negation,[status(cth)],[t34_funct_2])).
fof(c_0_8, plain, ![X13, X14, X15]:((((~v1_funct_2(X15,X13,X14)|X13=k1_relset_1(X13,X14,X15)|X14=k1_xboole_0|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))&(X13!=k1_relset_1(X13,X14,X15)|v1_funct_2(X15,X13,X14)|X14=k1_xboole_0|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14)))))&((~v1_funct_2(X15,X13,X14)|X13=k1_relset_1(X13,X14,X15)|X13!=k1_xboole_0|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))&(X13!=k1_relset_1(X13,X14,X15)|v1_funct_2(X15,X13,X14)|X13!=k1_xboole_0|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))))&((~v1_funct_2(X15,X13,X14)|X15=k1_xboole_0|X13=k1_xboole_0|X14!=k1_xboole_0|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))&(X15!=k1_xboole_0|v1_funct_2(X15,X13,X14)|X13=k1_xboole_0|X14!=k1_xboole_0|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
fof(c_0_9, negated_conjecture, ![X11, X12]:((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,esk1_0,esk2_0)&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))&((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,esk2_0,esk1_0)&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))))&((k2_relset_1(esk1_0,esk2_0,esk3_0)=esk2_0&(v2_funct_1(esk3_0)&(((r2_hidden(X12,esk1_0)|(~r2_hidden(X11,esk2_0)|k1_funct_1(esk4_0,X11)!=X12))&(k1_funct_1(esk3_0,X12)=X11|(~r2_hidden(X11,esk2_0)|k1_funct_1(esk4_0,X11)!=X12)))&((r2_hidden(X11,esk2_0)|(~r2_hidden(X12,esk1_0)|k1_funct_1(esk3_0,X12)!=X11))&(k1_funct_1(esk4_0,X11)=X12|(~r2_hidden(X12,esk1_0)|k1_funct_1(esk3_0,X12)!=X11))))))&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&esk4_0!=k2_funct_1(esk3_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).
fof(c_0_10, plain, ![X16, X17]:(~v1_relat_1(X16)|(~m1_subset_1(X17,k1_zfmisc_1(X16))|v1_relat_1(X17))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_11, plain, ![X18, X19]:v1_relat_1(k2_zfmisc_1(X18,X19)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_12, plain, ![X20, X21, X22]:(~m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(X20,X21)))|k1_relset_1(X20,X21,X22)=k1_relat_1(X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_13, plain, (X2=k1_relset_1(X2,X3,X1)|X3=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, negated_conjecture, (v1_funct_2(esk4_0,esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_17, plain, ![X23, X24, X25, X26, X27, X28]:(((k1_relat_1(X24)=k2_relat_1(X23)|X24!=k2_funct_1(X23)|(~v1_relat_1(X24)|~v1_funct_1(X24))|~v2_funct_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23)))&(((r2_hidden(X26,k1_relat_1(X23))|(~r2_hidden(X25,k2_relat_1(X23))|X26!=k1_funct_1(X24,X25))|X24!=k2_funct_1(X23)|(~v1_relat_1(X24)|~v1_funct_1(X24))|~v2_funct_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23)))&(X25=k1_funct_1(X23,X26)|(~r2_hidden(X25,k2_relat_1(X23))|X26!=k1_funct_1(X24,X25))|X24!=k2_funct_1(X23)|(~v1_relat_1(X24)|~v1_funct_1(X24))|~v2_funct_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23))))&((r2_hidden(X27,k2_relat_1(X23))|(~r2_hidden(X28,k1_relat_1(X23))|X27!=k1_funct_1(X23,X28))|X24!=k2_funct_1(X23)|(~v1_relat_1(X24)|~v1_funct_1(X24))|~v2_funct_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23)))&(X28=k1_funct_1(X24,X27)|(~r2_hidden(X28,k1_relat_1(X23))|X27!=k1_funct_1(X23,X28))|X24!=k2_funct_1(X23)|(~v1_relat_1(X24)|~v1_funct_1(X24))|~v2_funct_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23))))))&(((((r2_hidden(esk8_2(X23,X24),k1_relat_1(X23))|r2_hidden(esk5_2(X23,X24),k2_relat_1(X23))|k1_relat_1(X24)!=k2_relat_1(X23)|X24=k2_funct_1(X23)|(~v1_relat_1(X24)|~v1_funct_1(X24))|~v2_funct_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23)))&(esk7_2(X23,X24)=k1_funct_1(X23,esk8_2(X23,X24))|r2_hidden(esk5_2(X23,X24),k2_relat_1(X23))|k1_relat_1(X24)!=k2_relat_1(X23)|X24=k2_funct_1(X23)|(~v1_relat_1(X24)|~v1_funct_1(X24))|~v2_funct_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23))))&(~r2_hidden(esk7_2(X23,X24),k2_relat_1(X23))|esk8_2(X23,X24)!=k1_funct_1(X24,esk7_2(X23,X24))|r2_hidden(esk5_2(X23,X24),k2_relat_1(X23))|k1_relat_1(X24)!=k2_relat_1(X23)|X24=k2_funct_1(X23)|(~v1_relat_1(X24)|~v1_funct_1(X24))|~v2_funct_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23))))&(((r2_hidden(esk8_2(X23,X24),k1_relat_1(X23))|esk6_2(X23,X24)=k1_funct_1(X24,esk5_2(X23,X24))|k1_relat_1(X24)!=k2_relat_1(X23)|X24=k2_funct_1(X23)|(~v1_relat_1(X24)|~v1_funct_1(X24))|~v2_funct_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23)))&(esk7_2(X23,X24)=k1_funct_1(X23,esk8_2(X23,X24))|esk6_2(X23,X24)=k1_funct_1(X24,esk5_2(X23,X24))|k1_relat_1(X24)!=k2_relat_1(X23)|X24=k2_funct_1(X23)|(~v1_relat_1(X24)|~v1_funct_1(X24))|~v2_funct_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23))))&(~r2_hidden(esk7_2(X23,X24),k2_relat_1(X23))|esk8_2(X23,X24)!=k1_funct_1(X24,esk7_2(X23,X24))|esk6_2(X23,X24)=k1_funct_1(X24,esk5_2(X23,X24))|k1_relat_1(X24)!=k2_relat_1(X23)|X24=k2_funct_1(X23)|(~v1_relat_1(X24)|~v1_funct_1(X24))|~v2_funct_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23)))))&(((r2_hidden(esk8_2(X23,X24),k1_relat_1(X23))|(~r2_hidden(esk6_2(X23,X24),k1_relat_1(X23))|esk5_2(X23,X24)!=k1_funct_1(X23,esk6_2(X23,X24)))|k1_relat_1(X24)!=k2_relat_1(X23)|X24=k2_funct_1(X23)|(~v1_relat_1(X24)|~v1_funct_1(X24))|~v2_funct_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23)))&(esk7_2(X23,X24)=k1_funct_1(X23,esk8_2(X23,X24))|(~r2_hidden(esk6_2(X23,X24),k1_relat_1(X23))|esk5_2(X23,X24)!=k1_funct_1(X23,esk6_2(X23,X24)))|k1_relat_1(X24)!=k2_relat_1(X23)|X24=k2_funct_1(X23)|(~v1_relat_1(X24)|~v1_funct_1(X24))|~v2_funct_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23))))&(~r2_hidden(esk7_2(X23,X24),k2_relat_1(X23))|esk8_2(X23,X24)!=k1_funct_1(X24,esk7_2(X23,X24))|(~r2_hidden(esk6_2(X23,X24),k1_relat_1(X23))|esk5_2(X23,X24)!=k1_funct_1(X23,esk6_2(X23,X24)))|k1_relat_1(X24)!=k2_relat_1(X23)|X24=k2_funct_1(X23)|(~v1_relat_1(X24)|~v1_funct_1(X24))|~v2_funct_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t54_funct_1])])])])])])).
cnf(c_0_18, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, negated_conjecture, (k1_relset_1(esk2_0,esk1_0,esk4_0)=esk2_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])]), c_0_16])).
fof(c_0_22, plain, ![X33, X34, X35]:(~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(X33,X34)))|k2_relset_1(X33,X34,X35)=k2_relat_1(X35)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
cnf(c_0_23, plain, (esk7_2(X1,X2)=k1_funct_1(X1,esk8_2(X1,X2))|esk6_2(X1,X2)=k1_funct_1(X2,esk5_2(X1,X2))|X2=k2_funct_1(X1)|k1_relat_1(X2)!=k2_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, negated_conjecture, (v1_relat_1(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_14]), c_0_19])])).
cnf(c_0_25, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_26, negated_conjecture, (k1_relat_1(esk4_0)=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_14]), c_0_21])).
cnf(c_0_27, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_29, negated_conjecture, (k2_relset_1(esk1_0,esk2_0,esk3_0)=esk2_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_30, negated_conjecture, (k1_funct_1(esk4_0,esk5_2(X1,esk4_0))=esk6_2(X1,esk4_0)|k1_funct_1(X1,esk8_2(X1,esk4_0))=esk7_2(X1,esk4_0)|k2_funct_1(X1)=esk4_0|k2_relat_1(X1)!=esk2_0|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_31, negated_conjecture, (k2_relat_1(esk3_0)=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])).
cnf(c_0_32, negated_conjecture, (v1_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_28]), c_0_19])])).
cnf(c_0_33, negated_conjecture, (v2_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_34, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_35, negated_conjecture, (esk4_0!=k2_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_36, plain, (esk7_2(X1,X2)=k1_funct_1(X1,esk8_2(X1,X2))|r2_hidden(esk5_2(X1,X2),k2_relat_1(X1))|X2=k2_funct_1(X1)|k1_relat_1(X2)!=k2_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_37, negated_conjecture, (v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_38, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_39, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X2,esk2_0)|k1_funct_1(esk4_0,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_40, negated_conjecture, (k1_funct_1(esk3_0,esk8_2(esk3_0,esk4_0))=esk7_2(esk3_0,esk4_0)|k1_funct_1(esk4_0,esk5_2(esk3_0,esk4_0))=esk6_2(esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_41, negated_conjecture, (k1_funct_1(X1,esk8_2(X1,esk4_0))=esk7_2(X1,esk4_0)|k2_funct_1(X1)=esk4_0|r2_hidden(esk5_2(X1,esk4_0),k2_relat_1(X1))|k2_relat_1(X1)!=esk2_0|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_42, negated_conjecture, (k1_relset_1(esk1_0,esk2_0,esk3_0)=esk1_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_28]), c_0_37])]), c_0_38])).
cnf(c_0_43, negated_conjecture, (k1_funct_1(esk3_0,esk8_2(esk3_0,esk4_0))=esk7_2(esk3_0,esk4_0)|r2_hidden(X1,esk1_0)|esk6_2(esk3_0,esk4_0)!=X1|~r2_hidden(esk5_2(esk3_0,esk4_0),esk2_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_44, negated_conjecture, (k1_funct_1(esk3_0,esk8_2(esk3_0,esk4_0))=esk7_2(esk3_0,esk4_0)|r2_hidden(esk5_2(esk3_0,esk4_0),esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_31]), c_0_32]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_45, plain, (r2_hidden(esk8_2(X1,X2),k1_relat_1(X1))|r2_hidden(esk5_2(X1,X2),k2_relat_1(X1))|X2=k2_funct_1(X1)|k1_relat_1(X2)!=k2_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_46, negated_conjecture, (k1_funct_1(esk3_0,X1)=X2|~r2_hidden(X2,esk2_0)|k1_funct_1(esk4_0,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_47, plain, (esk7_2(X1,X2)=k1_funct_1(X1,esk8_2(X1,X2))|X2=k2_funct_1(X1)|~r2_hidden(esk6_2(X1,X2),k1_relat_1(X1))|esk5_2(X1,X2)!=k1_funct_1(X1,esk6_2(X1,X2))|k1_relat_1(X2)!=k2_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_48, negated_conjecture, (k1_relat_1(esk3_0)=esk1_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_28]), c_0_42])).
cnf(c_0_49, negated_conjecture, (k1_funct_1(esk3_0,esk8_2(esk3_0,esk4_0))=esk7_2(esk3_0,esk4_0)|r2_hidden(X1,esk1_0)|esk6_2(esk3_0,esk4_0)!=X1), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_50, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X2,esk1_0)|k1_funct_1(esk3_0,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_51, negated_conjecture, (k2_funct_1(X1)=esk4_0|r2_hidden(esk8_2(X1,esk4_0),k1_relat_1(X1))|r2_hidden(esk5_2(X1,esk4_0),k2_relat_1(X1))|k2_relat_1(X1)!=esk2_0|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_52, negated_conjecture, (k1_funct_1(esk3_0,k1_funct_1(esk4_0,X1))=X1|~r2_hidden(X1,esk2_0)), inference(er,[status(thm)],[c_0_46])).
cnf(c_0_53, negated_conjecture, (k1_funct_1(esk3_0,esk8_2(esk3_0,X1))=esk7_2(esk3_0,X1)|X1=k2_funct_1(esk3_0)|k1_funct_1(esk3_0,esk6_2(esk3_0,X1))!=esk5_2(esk3_0,X1)|k1_relat_1(X1)!=esk2_0|~v1_relat_1(X1)|~r2_hidden(esk6_2(esk3_0,X1),esk1_0)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_31]), c_0_32]), c_0_33]), c_0_34])])).
cnf(c_0_54, negated_conjecture, (k1_funct_1(esk3_0,esk8_2(esk3_0,esk4_0))=esk7_2(esk3_0,esk4_0)|r2_hidden(esk6_2(esk3_0,esk4_0),esk1_0)), inference(er,[status(thm)],[c_0_49])).
cnf(c_0_55, negated_conjecture, (k1_funct_1(esk4_0,X1)=X2|~r2_hidden(X2,esk1_0)|k1_funct_1(esk3_0,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_56, plain, (r2_hidden(esk8_2(X1,X2),k1_relat_1(X1))|esk6_2(X1,X2)=k1_funct_1(X2,esk5_2(X1,X2))|X2=k2_funct_1(X1)|k1_relat_1(X2)!=k2_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_57, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,X1),esk2_0)|~r2_hidden(X1,esk1_0)), inference(er,[status(thm)],[c_0_50])).
cnf(c_0_58, negated_conjecture, (r2_hidden(esk5_2(esk3_0,esk4_0),esk2_0)|r2_hidden(esk8_2(esk3_0,esk4_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_31]), c_0_48]), c_0_32]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_59, negated_conjecture, (k1_funct_1(esk3_0,k1_funct_1(esk4_0,esk5_2(esk3_0,esk4_0)))=esk5_2(esk3_0,esk4_0)|k1_funct_1(esk3_0,esk8_2(esk3_0,esk4_0))=esk7_2(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_52, c_0_44])).
cnf(c_0_60, negated_conjecture, (k1_funct_1(esk3_0,esk8_2(esk3_0,esk4_0))=esk7_2(esk3_0,esk4_0)|k1_funct_1(esk3_0,esk6_2(esk3_0,esk4_0))!=esk5_2(esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_26]), c_0_24]), c_0_25])]), c_0_35])).
cnf(c_0_61, negated_conjecture, (k1_funct_1(esk4_0,k1_funct_1(esk3_0,X1))=X1|~r2_hidden(X1,esk1_0)), inference(er,[status(thm)],[c_0_55])).
cnf(c_0_62, negated_conjecture, (k1_funct_1(esk4_0,esk5_2(X1,esk4_0))=esk6_2(X1,esk4_0)|k2_funct_1(X1)=esk4_0|r2_hidden(esk8_2(X1,esk4_0),k1_relat_1(X1))|k2_relat_1(X1)!=esk2_0|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_63, plain, (r2_hidden(esk5_2(X1,X2),k2_relat_1(X1))|X2=k2_funct_1(X1)|~r2_hidden(esk7_2(X1,X2),k2_relat_1(X1))|esk8_2(X1,X2)!=k1_funct_1(X2,esk7_2(X1,X2))|k1_relat_1(X2)!=k2_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_64, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk8_2(esk3_0,esk4_0)),esk2_0)|r2_hidden(esk5_2(esk3_0,esk4_0),esk2_0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_65, negated_conjecture, (k1_funct_1(esk3_0,esk8_2(esk3_0,esk4_0))=esk7_2(esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_40]), c_0_60])).
cnf(c_0_66, negated_conjecture, (k1_funct_1(esk4_0,k1_funct_1(esk3_0,esk8_2(esk3_0,esk4_0)))=esk8_2(esk3_0,esk4_0)|r2_hidden(esk5_2(esk3_0,esk4_0),esk2_0)), inference(spm,[status(thm)],[c_0_61, c_0_58])).
cnf(c_0_67, negated_conjecture, (k1_funct_1(esk4_0,esk5_2(esk3_0,esk4_0))=esk6_2(esk3_0,esk4_0)|r2_hidden(esk8_2(esk3_0,esk4_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_31]), c_0_48]), c_0_32]), c_0_33]), c_0_34])]), c_0_35])).
cnf(c_0_68, negated_conjecture, (X1=k2_funct_1(esk3_0)|r2_hidden(esk5_2(esk3_0,X1),esk2_0)|k1_funct_1(X1,esk7_2(esk3_0,X1))!=esk8_2(esk3_0,X1)|k1_relat_1(X1)!=esk2_0|~v1_relat_1(X1)|~r2_hidden(esk7_2(esk3_0,X1),esk2_0)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_31]), c_0_32]), c_0_33]), c_0_34])])).
cnf(c_0_69, negated_conjecture, (r2_hidden(esk5_2(esk3_0,esk4_0),esk2_0)|r2_hidden(esk7_2(esk3_0,esk4_0),esk2_0)), inference(rw,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_70, negated_conjecture, (k1_funct_1(esk4_0,esk7_2(esk3_0,esk4_0))=esk8_2(esk3_0,esk4_0)|r2_hidden(esk5_2(esk3_0,esk4_0),esk2_0)), inference(rw,[status(thm)],[c_0_66, c_0_65])).
cnf(c_0_71, plain, (esk6_2(X1,X2)=k1_funct_1(X2,esk5_2(X1,X2))|X2=k2_funct_1(X1)|~r2_hidden(esk7_2(X1,X2),k2_relat_1(X1))|esk8_2(X1,X2)!=k1_funct_1(X2,esk7_2(X1,X2))|k1_relat_1(X2)!=k2_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_72, negated_conjecture, (k1_funct_1(esk4_0,esk5_2(esk3_0,esk4_0))=esk6_2(esk3_0,esk4_0)|r2_hidden(k1_funct_1(esk3_0,esk8_2(esk3_0,esk4_0)),esk2_0)), inference(spm,[status(thm)],[c_0_57, c_0_67])).
cnf(c_0_73, negated_conjecture, (k1_funct_1(esk4_0,k1_funct_1(esk3_0,esk8_2(esk3_0,esk4_0)))=esk8_2(esk3_0,esk4_0)|k1_funct_1(esk4_0,esk5_2(esk3_0,esk4_0))=esk6_2(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_61, c_0_67])).
cnf(c_0_74, negated_conjecture, (r2_hidden(k1_funct_1(esk4_0,X1),esk1_0)|~r2_hidden(X1,esk2_0)), inference(er,[status(thm)],[c_0_39])).
cnf(c_0_75, negated_conjecture, (r2_hidden(esk5_2(esk3_0,esk4_0),esk2_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_26]), c_0_24]), c_0_25])]), c_0_35]), c_0_70])).
cnf(c_0_76, negated_conjecture, (k1_funct_1(X1,esk5_2(esk3_0,X1))=esk6_2(esk3_0,X1)|X1=k2_funct_1(esk3_0)|k1_funct_1(X1,esk7_2(esk3_0,X1))!=esk8_2(esk3_0,X1)|k1_relat_1(X1)!=esk2_0|~v1_relat_1(X1)|~r2_hidden(esk7_2(esk3_0,X1),esk2_0)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_31]), c_0_32]), c_0_33]), c_0_34])])).
cnf(c_0_77, negated_conjecture, (k1_funct_1(esk4_0,esk5_2(esk3_0,esk4_0))=esk6_2(esk3_0,esk4_0)|r2_hidden(esk7_2(esk3_0,esk4_0),esk2_0)), inference(rw,[status(thm)],[c_0_72, c_0_65])).
cnf(c_0_78, negated_conjecture, (k1_funct_1(esk4_0,esk5_2(esk3_0,esk4_0))=esk6_2(esk3_0,esk4_0)|k1_funct_1(esk4_0,esk7_2(esk3_0,esk4_0))=esk8_2(esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_73, c_0_65])).
cnf(c_0_79, plain, (r2_hidden(esk8_2(X1,X2),k1_relat_1(X1))|X2=k2_funct_1(X1)|~r2_hidden(esk6_2(X1,X2),k1_relat_1(X1))|esk5_2(X1,X2)!=k1_funct_1(X1,esk6_2(X1,X2))|k1_relat_1(X2)!=k2_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_80, negated_conjecture, (r2_hidden(k1_funct_1(esk4_0,esk5_2(esk3_0,esk4_0)),esk1_0)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_81, negated_conjecture, (k1_funct_1(esk4_0,esk5_2(esk3_0,esk4_0))=esk6_2(esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_26]), c_0_24]), c_0_25])]), c_0_35]), c_0_78])).
cnf(c_0_82, negated_conjecture, (k1_funct_1(esk3_0,k1_funct_1(esk4_0,esk5_2(esk3_0,esk4_0)))=esk5_2(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_52, c_0_75])).
cnf(c_0_83, negated_conjecture, (X1=k2_funct_1(esk3_0)|r2_hidden(esk8_2(esk3_0,X1),esk1_0)|k1_funct_1(esk3_0,esk6_2(esk3_0,X1))!=esk5_2(esk3_0,X1)|k1_relat_1(X1)!=esk2_0|~v1_relat_1(X1)|~r2_hidden(esk6_2(esk3_0,X1),esk1_0)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_48]), c_0_31]), c_0_32]), c_0_33]), c_0_34])])).
cnf(c_0_84, negated_conjecture, (r2_hidden(esk6_2(esk3_0,esk4_0),esk1_0)), inference(rw,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_85, negated_conjecture, (k1_funct_1(esk3_0,esk6_2(esk3_0,esk4_0))=esk5_2(esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_82, c_0_81])).
cnf(c_0_86, plain, (X2=k2_funct_1(X1)|~r2_hidden(esk7_2(X1,X2),k2_relat_1(X1))|esk8_2(X1,X2)!=k1_funct_1(X2,esk7_2(X1,X2))|~r2_hidden(esk6_2(X1,X2),k1_relat_1(X1))|esk5_2(X1,X2)!=k1_funct_1(X1,esk6_2(X1,X2))|k1_relat_1(X2)!=k2_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_87, negated_conjecture, (r2_hidden(esk8_2(esk3_0,esk4_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_26]), c_0_24]), c_0_25])]), c_0_35]), c_0_85])])).
cnf(c_0_88, negated_conjecture, (X1=k2_funct_1(esk3_0)|k1_funct_1(esk3_0,esk6_2(esk3_0,X1))!=esk5_2(esk3_0,X1)|k1_funct_1(X1,esk7_2(esk3_0,X1))!=esk8_2(esk3_0,X1)|k1_relat_1(X1)!=esk2_0|~v1_relat_1(X1)|~r2_hidden(esk6_2(esk3_0,X1),esk1_0)|~r2_hidden(esk7_2(esk3_0,X1),esk2_0)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_48]), c_0_31]), c_0_32]), c_0_31]), c_0_33]), c_0_34])])).
cnf(c_0_89, negated_conjecture, (k1_funct_1(esk4_0,esk7_2(esk3_0,esk4_0))=esk8_2(esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_87]), c_0_65])).
cnf(c_0_90, negated_conjecture, (r2_hidden(esk7_2(esk3_0,esk4_0),esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_87]), c_0_65])).
cnf(c_0_91, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_84]), c_0_26]), c_0_24]), c_0_25])]), c_0_35]), c_0_85]), c_0_89]), c_0_90])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 92
# Proof object clause steps            : 77
# Proof object formula steps           : 15
# Proof object conjectures             : 66
# Proof object clause conjectures      : 63
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 29
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 42
# Proof object simplifying inferences  : 113
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 39
# Processed clauses                    : 371
# ...of these trivial                  : 10
# ...subsumed                          : 114
# ...remaining for further processing  : 246
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 8
# Backward-rewritten                   : 96
# Generated clauses                    : 460
# ...of the previous two non-trivial   : 475
# Contextual simplify-reflections      : 10
# Paramodulations                      : 410
# Factorizations                       : 0
# Equation resolutions                 : 50
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 103
#    Positive orientable unit clauses  : 25
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 75
# Current number of unprocessed clauses: 10
# ...number of literals in the above   : 43
# Current number of archived formulas  : 0
# Current number of archived clauses   : 143
# Clause-clause subsumption calls (NU) : 2316
# Rec. Clause-clause subsumption calls : 1008
# Non-unit clause-clause subsumptions  : 115
# Unit Clause-clause subsumption calls : 67
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 14446

# -------------------------------------------------
# User time                : 0.014 s
# System time              : 0.007 s
# Total time               : 0.021 s
# Maximum resident set size: 28748 pages
