# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_wellord2, axiom, ![X1, X2]:(v1_relat_1(X2)=>(X2=k1_wellord2(X1)<=>(k3_relat_1(X2)=X1&![X3, X4]:((r2_hidden(X3,X1)&r2_hidden(X4,X1))=>(r2_hidden(k4_tarski(X3,X4),X2)<=>r1_tarski(X3,X4)))))), file('data/MPTP/proofs/5cac0c2cbd664c12b2b14c40a799075e.E_input', d1_wellord2)).
fof(dt_k1_wellord2, axiom, ![X1]:v1_relat_1(k1_wellord2(X1)), file('data/MPTP/proofs/5cac0c2cbd664c12b2b14c40a799075e.E_input', dt_k1_wellord2)).
fof(t11_wellord2, conjecture, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r4_wellord1(k1_wellord2(X1),k1_wellord2(X2))=>X1=X2))), file('data/MPTP/proofs/5cac0c2cbd664c12b2b14c40a799075e.E_input', t11_wellord2)).
fof(t57_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v2_wellord1(X1)=>![X2]:~((r2_hidden(X2,k3_relat_1(X1))&r4_wellord1(X1,k2_wellord1(X1,k1_wellord1(X1,X2))))))), file('data/MPTP/proofs/5cac0c2cbd664c12b2b14c40a799075e.E_input', t57_wellord1)).
fof(t10_wellord2, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)=>X1=k1_wellord1(k1_wellord2(X2),X1)))), file('data/MPTP/proofs/5cac0c2cbd664c12b2b14c40a799075e.E_input', t10_wellord2)).
fof(t7_wellord2, axiom, ![X1]:(v3_ordinal1(X1)=>v2_wellord1(k1_wellord2(X1))), file('data/MPTP/proofs/5cac0c2cbd664c12b2b14c40a799075e.E_input', t7_wellord2)).
fof(t26_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r1_ordinal1(X1,X2)|r2_hidden(X2,X1)))), file('data/MPTP/proofs/5cac0c2cbd664c12b2b14c40a799075e.E_input', t26_ordinal1)).
fof(t8_wellord2, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_wellord1(k1_wellord2(X2),X1)=k1_wellord2(X1)), file('data/MPTP/proofs/5cac0c2cbd664c12b2b14c40a799075e.E_input', t8_wellord2)).
fof(redefinition_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)<=>r1_tarski(X1,X2))), file('data/MPTP/proofs/5cac0c2cbd664c12b2b14c40a799075e.E_input', redefinition_r1_ordinal1)).
fof(t50_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r4_wellord1(X1,X2)=>r4_wellord1(X2,X1)))), file('data/MPTP/proofs/5cac0c2cbd664c12b2b14c40a799075e.E_input', t50_wellord1)).
fof(connectedness_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)|r1_ordinal1(X2,X1))), file('data/MPTP/proofs/5cac0c2cbd664c12b2b14c40a799075e.E_input', connectedness_r1_ordinal1)).
fof(t24_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>~((~(r2_hidden(X1,X2))&(~(X1=X2)&~(r2_hidden(X2,X1))))))), file('data/MPTP/proofs/5cac0c2cbd664c12b2b14c40a799075e.E_input', t24_ordinal1)).
fof(c_0_12, plain, ![X5, X6, X7, X8]:(((k3_relat_1(X6)=X5|X6!=k1_wellord2(X5)|~v1_relat_1(X6))&((~r2_hidden(k4_tarski(X7,X8),X6)|r1_tarski(X7,X8)|(~r2_hidden(X7,X5)|~r2_hidden(X8,X5))|X6!=k1_wellord2(X5)|~v1_relat_1(X6))&(~r1_tarski(X7,X8)|r2_hidden(k4_tarski(X7,X8),X6)|(~r2_hidden(X7,X5)|~r2_hidden(X8,X5))|X6!=k1_wellord2(X5)|~v1_relat_1(X6))))&(((r2_hidden(esk1_2(X5,X6),X5)|k3_relat_1(X6)!=X5|X6=k1_wellord2(X5)|~v1_relat_1(X6))&(r2_hidden(esk2_2(X5,X6),X5)|k3_relat_1(X6)!=X5|X6=k1_wellord2(X5)|~v1_relat_1(X6)))&((~r2_hidden(k4_tarski(esk1_2(X5,X6),esk2_2(X5,X6)),X6)|~r1_tarski(esk1_2(X5,X6),esk2_2(X5,X6))|k3_relat_1(X6)!=X5|X6=k1_wellord2(X5)|~v1_relat_1(X6))&(r2_hidden(k4_tarski(esk1_2(X5,X6),esk2_2(X5,X6)),X6)|r1_tarski(esk1_2(X5,X6),esk2_2(X5,X6))|k3_relat_1(X6)!=X5|X6=k1_wellord2(X5)|~v1_relat_1(X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_wellord2])])])])])).
fof(c_0_13, plain, ![X11]:v1_relat_1(k1_wellord2(X11)), inference(variable_rename,[status(thm)],[dt_k1_wellord2])).
fof(c_0_14, negated_conjecture, ~(![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r4_wellord1(k1_wellord2(X1),k1_wellord2(X2))=>X1=X2)))), inference(assume_negation,[status(cth)],[t11_wellord2])).
fof(c_0_15, plain, ![X14, X15]:(~v1_relat_1(X14)|(~v2_wellord1(X14)|(~r2_hidden(X15,k3_relat_1(X14))|~r4_wellord1(X14,k2_wellord1(X14,k1_wellord1(X14,X15)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t57_wellord1])])])).
fof(c_0_16, plain, ![X16, X17]:(~v3_ordinal1(X16)|(~v3_ordinal1(X17)|(~r2_hidden(X16,X17)|X16=k1_wellord1(k1_wellord2(X17),X16)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_wellord2])])])).
cnf(c_0_17, plain, (k3_relat_1(X1)=X2|X1!=k1_wellord2(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, plain, (v1_relat_1(k1_wellord2(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_19, plain, ![X18]:(~v3_ordinal1(X18)|v2_wellord1(k1_wellord2(X18))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_wellord2])])).
fof(c_0_20, plain, ![X19, X20]:(~v3_ordinal1(X19)|(~v3_ordinal1(X20)|(r1_ordinal1(X19,X20)|r2_hidden(X20,X19)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t26_ordinal1])])])).
fof(c_0_21, negated_conjecture, (v3_ordinal1(esk3_0)&(v3_ordinal1(esk4_0)&(r4_wellord1(k1_wellord2(esk3_0),k1_wellord2(esk4_0))&esk3_0!=esk4_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
cnf(c_0_22, plain, (~v1_relat_1(X1)|~v2_wellord1(X1)|~r2_hidden(X2,k3_relat_1(X1))|~r4_wellord1(X1,k2_wellord1(X1,k1_wellord1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (X1=k1_wellord1(k1_wellord2(X2),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (k3_relat_1(k1_wellord2(X1))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_17]), c_0_18])])).
cnf(c_0_25, plain, (v2_wellord1(k1_wellord2(X1))|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_26, plain, ![X21, X22]:(~r1_tarski(X21,X22)|k2_wellord1(k1_wellord2(X22),X21)=k1_wellord2(X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_wellord2])])).
cnf(c_0_27, plain, (r1_ordinal1(X1,X2)|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, negated_conjecture, (v3_ordinal1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, plain, (~r4_wellord1(k1_wellord2(X1),k2_wellord1(k1_wellord2(X1),X2))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X2,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_18])]), c_0_24])]), c_0_25])).
cnf(c_0_30, plain, (k2_wellord1(k1_wellord2(X2),X1)=k1_wellord2(X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_31, plain, ![X23, X24]:((~r1_ordinal1(X23,X24)|r1_tarski(X23,X24)|(~v3_ordinal1(X23)|~v3_ordinal1(X24)))&(~r1_tarski(X23,X24)|r1_ordinal1(X23,X24)|(~v3_ordinal1(X23)|~v3_ordinal1(X24)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r1_ordinal1])])])).
cnf(c_0_32, negated_conjecture, (r1_ordinal1(X1,esk3_0)|r2_hidden(esk3_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_33, negated_conjecture, (v3_ordinal1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_34, plain, ![X25, X26]:(~v1_relat_1(X25)|(~v1_relat_1(X26)|(~r4_wellord1(X25,X26)|r4_wellord1(X26,X25)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t50_wellord1])])])).
fof(c_0_35, plain, ![X27, X28]:(~v3_ordinal1(X27)|~v3_ordinal1(X28)|(r1_ordinal1(X27,X28)|r1_ordinal1(X28,X27))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[connectedness_r1_ordinal1])])).
fof(c_0_36, plain, ![X29, X30]:(~v3_ordinal1(X29)|(~v3_ordinal1(X30)|(r2_hidden(X29,X30)|(X29=X30|r2_hidden(X30,X29))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_ordinal1])])])])).
cnf(c_0_37, plain, (~r4_wellord1(k1_wellord2(X1),k1_wellord2(X2))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_tarski(X2,X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_38, negated_conjecture, (r4_wellord1(k1_wellord2(esk3_0),k1_wellord2(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_39, plain, (r1_tarski(X1,X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_40, negated_conjecture, (r1_ordinal1(esk4_0,esk3_0)|r2_hidden(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_41, plain, (r4_wellord1(X2,X1)|~v1_relat_1(X1)|~v1_relat_1(X2)|~r4_wellord1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42, negated_conjecture, (r1_ordinal1(X1,esk4_0)|r2_hidden(esk4_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_27, c_0_33])).
cnf(c_0_43, plain, (r1_ordinal1(X1,X2)|r1_ordinal1(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_44, plain, (r2_hidden(X1,X2)|X1=X2|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_45, negated_conjecture, (~r1_tarski(esk4_0,esk3_0)|~r2_hidden(esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_28]), c_0_33])])).
cnf(c_0_46, negated_conjecture, (r1_tarski(esk4_0,esk3_0)|r2_hidden(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_28]), c_0_33])])).
cnf(c_0_47, negated_conjecture, (r4_wellord1(k1_wellord2(esk4_0),k1_wellord2(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_38]), c_0_18]), c_0_18])])).
cnf(c_0_48, negated_conjecture, (r1_ordinal1(esk3_0,esk4_0)|r2_hidden(esk4_0,esk3_0)), inference(spm,[status(thm)],[c_0_42, c_0_28])).
cnf(c_0_49, negated_conjecture, (r1_ordinal1(X1,esk4_0)|r1_ordinal1(esk4_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_43, c_0_33])).
cnf(c_0_50, negated_conjecture, (X1=esk4_0|r2_hidden(X1,esk4_0)|r2_hidden(esk4_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_44, c_0_33])).
cnf(c_0_51, negated_conjecture, (esk3_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk3_0,esk4_0)|~r2_hidden(esk4_0,esk3_0)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_53, negated_conjecture, (~r1_tarski(esk3_0,esk4_0)|~r2_hidden(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_47]), c_0_33]), c_0_28])])).
cnf(c_0_54, negated_conjecture, (r1_tarski(esk3_0,esk4_0)|r2_hidden(esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_48]), c_0_33]), c_0_28])])).
cnf(c_0_55, negated_conjecture, (r1_ordinal1(esk4_0,esk3_0)|r1_ordinal1(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_49, c_0_28])).
cnf(c_0_56, negated_conjecture, (r2_hidden(esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_28]), c_0_51]), c_0_52])).
cnf(c_0_57, negated_conjecture, (r2_hidden(esk4_0,esk3_0)|~r2_hidden(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_58, negated_conjecture, (r1_ordinal1(esk4_0,esk3_0)|r1_tarski(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_55]), c_0_33]), c_0_28])])).
cnf(c_0_59, negated_conjecture, (~r1_tarski(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_56])])).
cnf(c_0_60, negated_conjecture, (r2_hidden(esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_56])])).
cnf(c_0_61, negated_conjecture, (r1_ordinal1(esk4_0,esk3_0)), inference(sr,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_62, negated_conjecture, (~r1_tarski(esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_60])])).
cnf(c_0_63, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_61]), c_0_28]), c_0_33])]), c_0_62]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 64
# Proof object clause steps            : 39
# Proof object formula steps           : 25
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 38
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 22
# Processed clauses                    : 95
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 93
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 8
# Backward-rewritten                   : 18
# Generated clauses                    : 76
# ...of the previous two non-trivial   : 72
# Contextual simplify-reflections      : 2
# Paramodulations                      : 65
# Factorizations                       : 0
# Equation resolutions                 : 10
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 44
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 28
# Current number of unprocessed clauses: 21
# ...number of literals in the above   : 269
# Current number of archived formulas  : 0
# Current number of archived clauses   : 49
# Clause-clause subsumption calls (NU) : 1784
# Rec. Clause-clause subsumption calls : 213
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 123
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 7
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4689

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.003 s
# Total time               : 0.013 s
# Maximum resident set size: 28632 pages
