# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t81_zfmisc_1, conjecture, ![X1, X2]:r1_tarski(k2_xboole_0(k1_zfmisc_1(X1),k1_zfmisc_1(X2)),k1_zfmisc_1(k2_xboole_0(X1,X2))), file('data/MPTP/proofs/5ce9a785049141c3b866da626e64c4df.E_input', t81_zfmisc_1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/5ce9a785049141c3b866da626e64c4df.E_input', l51_zfmisc_1)).
fof(t77_enumset1, axiom, ![X1, X2]:k2_enumset1(X1,X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/5ce9a785049141c3b866da626e64c4df.E_input', t77_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/5ce9a785049141c3b866da626e64c4df.E_input', t72_enumset1)).
fof(t86_enumset1, axiom, ![X1, X2, X3, X4, X5]:k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/5ce9a785049141c3b866da626e64c4df.E_input', t86_enumset1)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/5ce9a785049141c3b866da626e64c4df.E_input', t98_xboole_1)).
fof(t8_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X3,X2))=>r1_tarski(k2_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/5ce9a785049141c3b866da626e64c4df.E_input', t8_xboole_1)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/5ce9a785049141c3b866da626e64c4df.E_input', d4_xboole_0)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/5ce9a785049141c3b866da626e64c4df.E_input', t48_xboole_1)).
fof(t65_zfmisc_1, axiom, ![X1, X2]:(k4_xboole_0(X1,k1_tarski(X2))=X1<=>~(r2_hidden(X2,X1))), file('data/MPTP/proofs/5ce9a785049141c3b866da626e64c4df.E_input', t65_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/5ce9a785049141c3b866da626e64c4df.E_input', t69_enumset1)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/5ce9a785049141c3b866da626e64c4df.E_input', t41_xboole_1)).
fof(t4_boole, axiom, ![X1]:k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0, file('data/MPTP/proofs/5ce9a785049141c3b866da626e64c4df.E_input', t4_boole)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/5ce9a785049141c3b866da626e64c4df.E_input', t7_xboole_1)).
fof(t53_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(X1,k2_xboole_0(X2,X3))=k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,X3)), file('data/MPTP/proofs/5ce9a785049141c3b866da626e64c4df.E_input', t53_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/5ce9a785049141c3b866da626e64c4df.E_input', t3_boole)).
fof(t79_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>r1_tarski(k1_zfmisc_1(X1),k1_zfmisc_1(X2))), file('data/MPTP/proofs/5ce9a785049141c3b866da626e64c4df.E_input', t79_zfmisc_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/5ce9a785049141c3b866da626e64c4df.E_input', t7_xboole_0)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/5ce9a785049141c3b866da626e64c4df.E_input', l32_xboole_1)).
fof(c_0_19, negated_conjecture, ~(![X1, X2]:r1_tarski(k2_xboole_0(k1_zfmisc_1(X1),k1_zfmisc_1(X2)),k1_zfmisc_1(k2_xboole_0(X1,X2)))), inference(assume_negation,[status(cth)],[t81_zfmisc_1])).
fof(c_0_20, negated_conjecture, ~r1_tarski(k2_xboole_0(k1_zfmisc_1(esk1_0),k1_zfmisc_1(esk2_0)),k1_zfmisc_1(k2_xboole_0(esk1_0,esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])).
fof(c_0_21, plain, ![X8, X9]:k3_tarski(k2_tarski(X8,X9))=k2_xboole_0(X8,X9), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_22, plain, ![X10, X11]:k2_enumset1(X10,X10,X10,X11)=k2_tarski(X10,X11), inference(variable_rename,[status(thm)],[t77_enumset1])).
fof(c_0_23, plain, ![X12, X13, X14, X15]:k3_enumset1(X12,X12,X13,X14,X15)=k2_enumset1(X12,X13,X14,X15), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_24, plain, ![X16, X17, X18, X19, X20]:k6_enumset1(X16,X16,X16,X16,X17,X18,X19,X20)=k3_enumset1(X16,X17,X18,X19,X20), inference(variable_rename,[status(thm)],[t86_enumset1])).
fof(c_0_25, plain, ![X21, X22]:k2_xboole_0(X21,X22)=k5_xboole_0(X21,k4_xboole_0(X22,X21)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_26, plain, ![X23, X24, X25]:(~r1_tarski(X23,X24)|~r1_tarski(X25,X24)|r1_tarski(k2_xboole_0(X23,X25),X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_xboole_1])])).
fof(c_0_27, plain, ![X26, X27, X28, X29, X30, X31, X32, X33]:((((r2_hidden(X29,X26)|~r2_hidden(X29,X28)|X28!=k3_xboole_0(X26,X27))&(r2_hidden(X29,X27)|~r2_hidden(X29,X28)|X28!=k3_xboole_0(X26,X27)))&(~r2_hidden(X30,X26)|~r2_hidden(X30,X27)|r2_hidden(X30,X28)|X28!=k3_xboole_0(X26,X27)))&((~r2_hidden(esk3_3(X31,X32,X33),X33)|(~r2_hidden(esk3_3(X31,X32,X33),X31)|~r2_hidden(esk3_3(X31,X32,X33),X32))|X33=k3_xboole_0(X31,X32))&((r2_hidden(esk3_3(X31,X32,X33),X31)|r2_hidden(esk3_3(X31,X32,X33),X33)|X33=k3_xboole_0(X31,X32))&(r2_hidden(esk3_3(X31,X32,X33),X32)|r2_hidden(esk3_3(X31,X32,X33),X33)|X33=k3_xboole_0(X31,X32))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_28, plain, ![X35, X36]:k4_xboole_0(X35,k4_xboole_0(X35,X36))=k3_xboole_0(X35,X36), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_29, plain, ![X37, X38]:((k4_xboole_0(X37,k1_tarski(X38))!=X37|~r2_hidden(X38,X37))&(r2_hidden(X38,X37)|k4_xboole_0(X37,k1_tarski(X38))=X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t65_zfmisc_1])])])).
fof(c_0_30, plain, ![X39]:k2_tarski(X39,X39)=k1_tarski(X39), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_31, negated_conjecture, (~r1_tarski(k2_xboole_0(k1_zfmisc_1(esk1_0),k1_zfmisc_1(esk2_0)),k1_zfmisc_1(k2_xboole_0(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_32, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_33, plain, (k2_enumset1(X1,X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_34, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_35, plain, (k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_36, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_37, plain, (r1_tarski(k2_xboole_0(X1,X3),X2)|~r1_tarski(X1,X2)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_38, plain, ![X40, X41, X42]:k4_xboole_0(k4_xboole_0(X40,X41),X42)=k4_xboole_0(X40,k2_xboole_0(X41,X42)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
cnf(c_0_39, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_40, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_41, plain, (k4_xboole_0(X1,k1_tarski(X2))!=X1|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_42, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_43, plain, ![X43]:k4_xboole_0(k1_xboole_0,X43)=k1_xboole_0, inference(variable_rename,[status(thm)],[t4_boole])).
cnf(c_0_44, negated_conjecture, (~r1_tarski(k3_tarski(k6_enumset1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(esk1_0),k1_zfmisc_1(esk1_0),k1_zfmisc_1(esk1_0),k1_zfmisc_1(esk1_0),k1_zfmisc_1(esk1_0),k1_zfmisc_1(esk1_0),k1_zfmisc_1(esk2_0))),k1_zfmisc_1(k3_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_32]), c_0_32]), c_0_33]), c_0_33]), c_0_34]), c_0_34]), c_0_35]), c_0_35])).
cnf(c_0_45, plain, (k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_32]), c_0_33]), c_0_34]), c_0_35])).
cnf(c_0_46, plain, (r1_tarski(k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X3)),X2)|~r1_tarski(X3,X2)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_32]), c_0_33]), c_0_34]), c_0_35])).
fof(c_0_47, plain, ![X44, X45]:r1_tarski(X44,k2_xboole_0(X44,X45)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_48, plain, ![X46, X47, X48]:k4_xboole_0(X46,k2_xboole_0(X47,X48))=k3_xboole_0(k4_xboole_0(X46,X47),k4_xboole_0(X46,X48)), inference(variable_rename,[status(thm)],[t53_xboole_1])).
cnf(c_0_49, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, plain, (r2_hidden(X1,X2)|X3!=k4_xboole_0(X4,k4_xboole_0(X4,X2))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_39, c_0_40])).
fof(c_0_51, plain, ![X49]:k4_xboole_0(X49,k1_xboole_0)=X49, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_52, plain, (k4_xboole_0(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2))!=X1|~r2_hidden(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_42]), c_0_33]), c_0_34]), c_0_35])).
cnf(c_0_53, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_54, negated_conjecture, (~r1_tarski(k5_xboole_0(k1_zfmisc_1(esk1_0),k4_xboole_0(k1_zfmisc_1(esk2_0),k1_zfmisc_1(esk1_0))),k1_zfmisc_1(k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_45]), c_0_45])).
cnf(c_0_55, plain, (r1_tarski(k5_xboole_0(X1,k4_xboole_0(X2,X1)),X3)|~r1_tarski(X2,X3)|~r1_tarski(X1,X3)), inference(rw,[status(thm)],[c_0_46, c_0_45])).
fof(c_0_56, plain, ![X50, X51]:(~r1_tarski(X50,X51)|r1_tarski(k1_zfmisc_1(X50),k1_zfmisc_1(X51))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t79_zfmisc_1])])).
cnf(c_0_57, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_58, plain, (k4_xboole_0(X1,k2_xboole_0(X2,X3))=k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k3_tarski(k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_32]), c_0_33]), c_0_34]), c_0_35])).
cnf(c_0_60, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X3,k4_xboole_0(X3,X2)))), inference(er,[status(thm)],[c_0_50])).
cnf(c_0_61, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_62, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
fof(c_0_63, plain, ![X52]:(X52=k1_xboole_0|r2_hidden(esk4_1(X52),X52)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_64, negated_conjecture, (~r1_tarski(k1_zfmisc_1(esk2_0),k1_zfmisc_1(k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0))))|~r1_tarski(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0))))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_65, plain, (r1_tarski(k1_zfmisc_1(X1),k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_66, plain, (r1_tarski(X1,k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_32]), c_0_33]), c_0_34]), c_0_35])).
cnf(c_0_67, plain, (k4_xboole_0(X1,k3_tarski(k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)))=k4_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_32]), c_0_33]), c_0_40]), c_0_34]), c_0_35])).
cnf(c_0_68, plain, (k4_xboole_0(X1,k5_xboole_0(X2,k4_xboole_0(X3,X2)))=k4_xboole_0(k4_xboole_0(X1,X2),X3)), inference(rw,[status(thm)],[c_0_59, c_0_45])).
cnf(c_0_69, plain, (~r2_hidden(X1,k4_xboole_0(X2,X2))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62])).
cnf(c_0_70, plain, (X1=k1_xboole_0|r2_hidden(esk4_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_71, negated_conjecture, (~r1_tarski(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0))))|~r1_tarski(esk2_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_72, plain, (r1_tarski(X1,k5_xboole_0(X1,k4_xboole_0(X2,X1)))), inference(rw,[status(thm)],[c_0_66, c_0_45])).
fof(c_0_73, plain, ![X54, X55]:((k4_xboole_0(X54,X55)!=k1_xboole_0|r1_tarski(X54,X55))&(~r1_tarski(X54,X55)|k4_xboole_0(X54,X55)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_74, plain, (k4_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,X3)))=k4_xboole_0(k4_xboole_0(X1,X2),X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_45]), c_0_68])).
cnf(c_0_75, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_76, negated_conjecture, (~r1_tarski(esk2_0,k5_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_65]), c_0_72])])).
cnf(c_0_77, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_78, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_61]), c_0_75])).
cnf(c_0_79, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_68]), c_0_78])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 80
# Proof object clause steps            : 41
# Proof object formula steps           : 39
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 50
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 6
# Initial clauses in saturation        : 20
# Processed clauses                    : 147
# ...of these trivial                  : 6
# ...subsumed                          : 53
# ...remaining for further processing  : 88
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 5
# Generated clauses                    : 626
# ...of the previous two non-trivial   : 506
# Contextual simplify-reflections      : 0
# Paramodulations                      : 618
# Factorizations                       : 4
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 59
#    Positive orientable unit clauses  : 18
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 35
# Current number of unprocessed clauses: 394
# ...number of literals in the above   : 1071
# Current number of archived formulas  : 0
# Current number of archived clauses   : 32
# Clause-clause subsumption calls (NU) : 289
# Rec. Clause-clause subsumption calls : 242
# Non-unit clause-clause subsumptions  : 21
# Unit Clause-clause subsumption calls : 39
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 31
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8272

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.003 s
# Total time               : 0.015 s
# Maximum resident set size: 28700 pages
