# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t83_mcart_1, conjecture, ![X1, X2, X3, X4, X5]:~((r2_hidden(X1,k4_zfmisc_1(X2,X3,X4,X5))&![X6, X7, X8, X9]:~((r2_hidden(X6,X2)&(r2_hidden(X7,X3)&(r2_hidden(X8,X4)&(r2_hidden(X9,X5)&X1=k4_mcart_1(X6,X7,X8,X9)))))))), file('data/MPTP/proofs/5d07e1f191a1458493bd733a76248ea5.E_input', t83_mcart_1)).
fof(d4_mcart_1, axiom, ![X1, X2, X3, X4]:k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4), file('data/MPTP/proofs/5d07e1f191a1458493bd733a76248ea5.E_input', d4_mcart_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/5d07e1f191a1458493bd733a76248ea5.E_input', d3_mcart_1)).
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/5d07e1f191a1458493bd733a76248ea5.E_input', d2_zfmisc_1)).
fof(d4_zfmisc_1, axiom, ![X1, X2, X3, X4]:k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4), file('data/MPTP/proofs/5d07e1f191a1458493bd733a76248ea5.E_input', d4_zfmisc_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/5d07e1f191a1458493bd733a76248ea5.E_input', d3_zfmisc_1)).
fof(c_0_6, negated_conjecture, ~(![X1, X2, X3, X4, X5]:~((r2_hidden(X1,k4_zfmisc_1(X2,X3,X4,X5))&![X6, X7, X8, X9]:~((r2_hidden(X6,X2)&(r2_hidden(X7,X3)&(r2_hidden(X8,X4)&(r2_hidden(X9,X5)&X1=k4_mcart_1(X6,X7,X8,X9))))))))), inference(assume_negation,[status(cth)],[t83_mcart_1])).
fof(c_0_7, plain, ![X26, X27, X28, X29]:k4_mcart_1(X26,X27,X28,X29)=k4_tarski(k3_mcart_1(X26,X27,X28),X29), inference(variable_rename,[status(thm)],[d4_mcart_1])).
fof(c_0_8, plain, ![X23, X24, X25]:k3_mcart_1(X23,X24,X25)=k4_tarski(k4_tarski(X23,X24),X25), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_9, negated_conjecture, ![X15, X16, X17, X18]:(r2_hidden(esk1_0,k4_zfmisc_1(esk2_0,esk3_0,esk4_0,esk5_0))&(~r2_hidden(X15,esk2_0)|(~r2_hidden(X16,esk3_0)|(~r2_hidden(X17,esk4_0)|(~r2_hidden(X18,esk5_0)|esk1_0!=k4_mcart_1(X15,X16,X17,X18)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
cnf(c_0_10, plain, (k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_12, plain, ![X33, X34, X35, X36, X39, X40, X41, X42, X43, X44, X46, X47]:((((r2_hidden(esk6_4(X33,X34,X35,X36),X33)|~r2_hidden(X36,X35)|X35!=k2_zfmisc_1(X33,X34))&((r2_hidden(esk7_4(X33,X34,X35,X36),X34)|~r2_hidden(X36,X35)|X35!=k2_zfmisc_1(X33,X34))&(X36=k4_tarski(esk6_4(X33,X34,X35,X36),esk7_4(X33,X34,X35,X36))|~r2_hidden(X36,X35)|X35!=k2_zfmisc_1(X33,X34))))&(~r2_hidden(X40,X33)|(~r2_hidden(X41,X34)|X39!=k4_tarski(X40,X41))|r2_hidden(X39,X35)|X35!=k2_zfmisc_1(X33,X34)))&((~r2_hidden(esk8_3(X42,X43,X44),X44)|(~r2_hidden(X46,X42)|(~r2_hidden(X47,X43)|esk8_3(X42,X43,X44)!=k4_tarski(X46,X47)))|X44=k2_zfmisc_1(X42,X43))&((r2_hidden(esk9_3(X42,X43,X44),X42)|r2_hidden(esk8_3(X42,X43,X44),X44)|X44=k2_zfmisc_1(X42,X43))&((r2_hidden(esk10_3(X42,X43,X44),X43)|r2_hidden(esk8_3(X42,X43,X44),X44)|X44=k2_zfmisc_1(X42,X43))&(esk8_3(X42,X43,X44)=k4_tarski(esk9_3(X42,X43,X44),esk10_3(X42,X43,X44))|r2_hidden(esk8_3(X42,X43,X44),X44)|X44=k2_zfmisc_1(X42,X43)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
cnf(c_0_13, negated_conjecture, (~r2_hidden(X1,esk2_0)|~r2_hidden(X2,esk3_0)|~r2_hidden(X3,esk4_0)|~r2_hidden(X4,esk5_0)|esk1_0!=k4_mcart_1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15, plain, (X1=k4_tarski(esk6_4(X2,X3,X4,X1),esk7_4(X2,X3,X4,X1))|~r2_hidden(X1,X4)|X4!=k2_zfmisc_1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, negated_conjecture, (esk1_0!=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)|~r2_hidden(X4,esk5_0)|~r2_hidden(X3,esk4_0)|~r2_hidden(X2,esk3_0)|~r2_hidden(X1,esk2_0)), inference(rw,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_17, plain, (k4_tarski(esk6_4(X1,X2,k2_zfmisc_1(X1,X2),X3),esk7_4(X1,X2,k2_zfmisc_1(X1,X2),X3))=X3|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_15])).
cnf(c_0_18, plain, (r2_hidden(esk7_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, negated_conjecture, (k4_tarski(k4_tarski(X1,X2),X3)!=esk1_0|~r2_hidden(esk7_4(X4,X5,k2_zfmisc_1(X4,X5),X1),esk3_0)|~r2_hidden(esk6_4(X4,X5,k2_zfmisc_1(X4,X5),X1),esk2_0)|~r2_hidden(X1,k2_zfmisc_1(X4,X5))|~r2_hidden(X3,esk5_0)|~r2_hidden(X2,esk4_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_20, plain, (r2_hidden(esk7_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X2)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_18])).
cnf(c_0_21, plain, (r2_hidden(esk6_4(X1,X2,X3,X4),X1)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22, negated_conjecture, (k4_tarski(k4_tarski(X1,X2),X3)!=esk1_0|~r2_hidden(esk6_4(X4,esk3_0,k2_zfmisc_1(X4,esk3_0),X1),esk2_0)|~r2_hidden(X1,k2_zfmisc_1(X4,esk3_0))|~r2_hidden(X3,esk5_0)|~r2_hidden(X2,esk4_0)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_23, plain, (r2_hidden(esk6_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X1)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_21])).
cnf(c_0_24, negated_conjecture, (k4_tarski(k4_tarski(X1,X2),X3)!=esk1_0|~r2_hidden(X1,k2_zfmisc_1(esk2_0,esk3_0))|~r2_hidden(X3,esk5_0)|~r2_hidden(X2,esk4_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_25, negated_conjecture, (k4_tarski(X1,X2)!=esk1_0|~r2_hidden(esk6_4(X3,X4,k2_zfmisc_1(X3,X4),X1),k2_zfmisc_1(esk2_0,esk3_0))|~r2_hidden(esk7_4(X3,X4,k2_zfmisc_1(X3,X4),X1),esk4_0)|~r2_hidden(X1,k2_zfmisc_1(X3,X4))|~r2_hidden(X2,esk5_0)), inference(spm,[status(thm)],[c_0_24, c_0_17])).
cnf(c_0_26, negated_conjecture, (k4_tarski(X1,X2)!=esk1_0|~r2_hidden(esk7_4(k2_zfmisc_1(esk2_0,esk3_0),X3,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),X3),X1),esk4_0)|~r2_hidden(X1,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),X3))|~r2_hidden(X2,esk5_0)), inference(spm,[status(thm)],[c_0_25, c_0_23])).
fof(c_0_27, plain, ![X19, X20, X21, X22]:k4_zfmisc_1(X19,X20,X21,X22)=k2_zfmisc_1(k3_zfmisc_1(X19,X20,X21),X22), inference(variable_rename,[status(thm)],[d4_zfmisc_1])).
fof(c_0_28, plain, ![X30, X31, X32]:k3_zfmisc_1(X30,X31,X32)=k2_zfmisc_1(k2_zfmisc_1(X30,X31),X32), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
cnf(c_0_29, negated_conjecture, (k4_tarski(X1,X2)!=esk1_0|~r2_hidden(X1,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0))|~r2_hidden(X2,esk5_0)), inference(spm,[status(thm)],[c_0_26, c_0_20])).
cnf(c_0_30, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_31, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_32, negated_conjecture, (k4_tarski(esk6_4(k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0),X1,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0),X1),X2),X3)!=esk1_0|~r2_hidden(X2,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0),X1))|~r2_hidden(X3,esk5_0)), inference(spm,[status(thm)],[c_0_29, c_0_23])).
cnf(c_0_33, negated_conjecture, (r2_hidden(esk1_0,k4_zfmisc_1(esk2_0,esk3_0,esk4_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_34, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_35, negated_conjecture, (~r2_hidden(esk7_4(k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0),X1,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0),X1),esk1_0),esk5_0)|~r2_hidden(esk1_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0),X1))), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_17])])).
cnf(c_0_36, negated_conjecture, (r2_hidden(esk1_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0),esk5_0))), inference(rw,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_37, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_20]), c_0_36])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 38
# Proof object clause steps            : 25
# Proof object formula steps           : 13
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 10
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 14
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 10
# Processed clauses                    : 41
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 40
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 45
# ...of the previous two non-trivial   : 42
# Contextual simplify-reflections      : 0
# Paramodulations                      : 40
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 26
#    Positive orientable unit clauses  : 1
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 25
# Current number of unprocessed clauses: 21
# ...number of literals in the above   : 114
# Current number of archived formulas  : 0
# Current number of archived clauses   : 14
# Clause-clause subsumption calls (NU) : 37
# Rec. Clause-clause subsumption calls : 13
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2323

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.003 s
# Total time               : 0.009 s
# Maximum resident set size: 28632 pages
