# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t71_mcart_1, conjecture, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X8))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k7_mcart_1(X1,X2,X3,X5)))))), file('data/MPTP/proofs/5d5aa40200514d709c39d49ba4dcef09.E_input', t71_mcart_1)).
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/5d5aa40200514d709c39d49ba4dcef09.E_input', d2_zfmisc_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/5d5aa40200514d709c39d49ba4dcef09.E_input', d3_zfmisc_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/5d5aa40200514d709c39d49ba4dcef09.E_input', d1_xboole_0)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/5d5aa40200514d709c39d49ba4dcef09.E_input', d2_subset_1)).
fof(t34_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4, X5, X6]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k4_mcart_1(X3,X4,X5,X6))))))), file('data/MPTP/proofs/5d5aa40200514d709c39d49ba4dcef09.E_input', t34_mcart_1)).
fof(t72_zfmisc_1, axiom, ![X1, X2, X3]:(k4_xboole_0(k2_tarski(X1,X2),X3)=k2_tarski(X1,X2)<=>(~(r2_hidden(X1,X3))&~(r2_hidden(X2,X3)))), file('data/MPTP/proofs/5d5aa40200514d709c39d49ba4dcef09.E_input', t72_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/5d5aa40200514d709c39d49ba4dcef09.E_input', t70_enumset1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/5d5aa40200514d709c39d49ba4dcef09.E_input', t100_xboole_1)).
fof(t46_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0, file('data/MPTP/proofs/5d5aa40200514d709c39d49ba4dcef09.E_input', t46_xboole_1)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/5d5aa40200514d709c39d49ba4dcef09.E_input', d2_tarski)).
fof(t35_mcart_1, axiom, ![X1, X2, X3]:((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&~(X3=k1_xboole_0)))<=>~(k3_zfmisc_1(X1,X2,X3)=k1_xboole_0)), file('data/MPTP/proofs/5d5aa40200514d709c39d49ba4dcef09.E_input', t35_mcart_1)).
fof(t95_enumset1, axiom, ![X1, X2]:k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/5d5aa40200514d709c39d49ba4dcef09.E_input', t95_enumset1)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/5d5aa40200514d709c39d49ba4dcef09.E_input', t1_boole)).
fof(t10_mcart_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(r2_hidden(k1_mcart_1(X1),X2)&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/5d5aa40200514d709c39d49ba4dcef09.E_input', t10_mcart_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/5d5aa40200514d709c39d49ba4dcef09.E_input', d3_mcart_1)).
fof(t23_mcart_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r2_hidden(X1,X2)=>X1=k4_tarski(k1_mcart_1(X1),k2_mcart_1(X1)))), file('data/MPTP/proofs/5d5aa40200514d709c39d49ba4dcef09.E_input', t23_mcart_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/5d5aa40200514d709c39d49ba4dcef09.E_input', fc6_relat_1)).
fof(t50_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))&(k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))&k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4))))))))), file('data/MPTP/proofs/5d5aa40200514d709c39d49ba4dcef09.E_input', t50_mcart_1)).
fof(c_0_19, negated_conjecture, ~(![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X8))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k7_mcart_1(X1,X2,X3,X5))))))), inference(assume_negation,[status(cth)],[t71_mcart_1])).
fof(c_0_20, plain, ![X37, X38, X39, X40, X43, X44, X45, X46, X47, X48, X50, X51]:((((r2_hidden(esk7_4(X37,X38,X39,X40),X37)|~r2_hidden(X40,X39)|X39!=k2_zfmisc_1(X37,X38))&((r2_hidden(esk8_4(X37,X38,X39,X40),X38)|~r2_hidden(X40,X39)|X39!=k2_zfmisc_1(X37,X38))&(X40=k4_tarski(esk7_4(X37,X38,X39,X40),esk8_4(X37,X38,X39,X40))|~r2_hidden(X40,X39)|X39!=k2_zfmisc_1(X37,X38))))&(~r2_hidden(X44,X37)|(~r2_hidden(X45,X38)|X43!=k4_tarski(X44,X45))|r2_hidden(X43,X39)|X39!=k2_zfmisc_1(X37,X38)))&((~r2_hidden(esk9_3(X46,X47,X48),X48)|(~r2_hidden(X50,X46)|(~r2_hidden(X51,X47)|esk9_3(X46,X47,X48)!=k4_tarski(X50,X51)))|X48=k2_zfmisc_1(X46,X47))&((r2_hidden(esk10_3(X46,X47,X48),X46)|r2_hidden(esk9_3(X46,X47,X48),X48)|X48=k2_zfmisc_1(X46,X47))&((r2_hidden(esk11_3(X46,X47,X48),X47)|r2_hidden(esk9_3(X46,X47,X48),X48)|X48=k2_zfmisc_1(X46,X47))&(esk9_3(X46,X47,X48)=k4_tarski(esk10_3(X46,X47,X48),esk11_3(X46,X47,X48))|r2_hidden(esk9_3(X46,X47,X48),X48)|X48=k2_zfmisc_1(X46,X47)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
fof(c_0_21, negated_conjecture, ![X20, X21, X22]:(m1_subset_1(esk5_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))&((~m1_subset_1(X20,esk1_0)|(~m1_subset_1(X21,esk2_0)|(~m1_subset_1(X22,esk3_0)|(esk5_0!=k3_mcart_1(X20,X21,X22)|esk4_0=X22))))&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&esk4_0!=k7_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])).
fof(c_0_22, plain, ![X12, X13, X14]:k3_zfmisc_1(X12,X13,X14)=k2_zfmisc_1(k2_zfmisc_1(X12,X13),X14), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_23, plain, ![X54, X55, X56]:((~v1_xboole_0(X54)|~r2_hidden(X55,X54))&(r2_hidden(esk12_1(X56),X56)|v1_xboole_0(X56))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_24, plain, (r2_hidden(X5,X6)|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)|X5!=k4_tarski(X1,X3)|X6!=k2_zfmisc_1(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_25, plain, ![X58, X59]:(((~m1_subset_1(X59,X58)|r2_hidden(X59,X58)|v1_xboole_0(X58))&(~r2_hidden(X59,X58)|m1_subset_1(X59,X58)|v1_xboole_0(X58)))&((~m1_subset_1(X59,X58)|v1_xboole_0(X59)|~v1_xboole_0(X58))&(~v1_xboole_0(X59)|m1_subset_1(X59,X58)|~v1_xboole_0(X58)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
cnf(c_0_26, negated_conjecture, (m1_subset_1(esk5_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_24])])).
cnf(c_0_30, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(rw,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32, plain, (r2_hidden(esk7_4(X1,X2,X3,X4),X1)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_33, plain, (~v1_xboole_0(k2_zfmisc_1(X1,X2))|~r2_hidden(X3,X2)|~r2_hidden(X4,X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_35, plain, (r2_hidden(esk7_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X1)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_32])).
fof(c_0_36, plain, ![X60, X62, X63, X64, X65]:((r2_hidden(esk13_1(X60),X60)|X60=k1_xboole_0)&((~r2_hidden(X62,X60)|esk13_1(X60)!=k4_mcart_1(X62,X63,X64,X65)|X60=k1_xboole_0)&(~r2_hidden(X63,X60)|esk13_1(X60)!=k4_mcart_1(X62,X63,X64,X65)|X60=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_mcart_1])])])])])])])).
cnf(c_0_37, negated_conjecture, (r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|~r2_hidden(X1,k2_zfmisc_1(esk1_0,esk2_0))|~r2_hidden(X2,esk3_0)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_38, plain, (r2_hidden(esk12_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_39, plain, ![X66, X67, X68]:(((~r2_hidden(X66,X68)|k4_xboole_0(k2_tarski(X66,X67),X68)!=k2_tarski(X66,X67))&(~r2_hidden(X67,X68)|k4_xboole_0(k2_tarski(X66,X67),X68)!=k2_tarski(X66,X67)))&(r2_hidden(X66,X68)|r2_hidden(X67,X68)|k4_xboole_0(k2_tarski(X66,X67),X68)=k2_tarski(X66,X67))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t72_zfmisc_1])])])])).
fof(c_0_40, plain, ![X32, X33]:k1_enumset1(X32,X32,X33)=k2_tarski(X32,X33), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_41, plain, ![X69, X70]:k4_xboole_0(X69,X70)=k5_xboole_0(X69,k3_xboole_0(X69,X70)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_42, plain, ![X71, X72]:k4_xboole_0(X71,k2_xboole_0(X71,X72))=k1_xboole_0, inference(variable_rename,[status(thm)],[t46_xboole_1])).
fof(c_0_43, plain, ![X23, X24, X25, X26, X27, X28, X29, X30]:(((~r2_hidden(X26,X25)|(X26=X23|X26=X24)|X25!=k2_tarski(X23,X24))&((X27!=X23|r2_hidden(X27,X25)|X25!=k2_tarski(X23,X24))&(X27!=X24|r2_hidden(X27,X25)|X25!=k2_tarski(X23,X24))))&(((esk6_3(X28,X29,X30)!=X28|~r2_hidden(esk6_3(X28,X29,X30),X30)|X30=k2_tarski(X28,X29))&(esk6_3(X28,X29,X30)!=X29|~r2_hidden(esk6_3(X28,X29,X30),X30)|X30=k2_tarski(X28,X29)))&(r2_hidden(esk6_3(X28,X29,X30),X30)|(esk6_3(X28,X29,X30)=X28|esk6_3(X28,X29,X30)=X29)|X30=k2_tarski(X28,X29)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_44, plain, ![X9, X10, X11]:((X9=k1_xboole_0|(X10=k1_xboole_0|X11=k1_xboole_0)|k3_zfmisc_1(X9,X10,X11)!=k1_xboole_0)&((X9!=k1_xboole_0|k3_zfmisc_1(X9,X10,X11)=k1_xboole_0)&((X10!=k1_xboole_0|k3_zfmisc_1(X9,X10,X11)=k1_xboole_0)&(X11!=k1_xboole_0|k3_zfmisc_1(X9,X10,X11)=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t35_mcart_1])])])])).
cnf(c_0_45, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k2_zfmisc_1(X1,X3))), inference(spm,[status(thm)],[c_0_28, c_0_35])).
cnf(c_0_46, plain, (r2_hidden(esk13_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_47, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0))|r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_48, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_49, plain, (~r2_hidden(X1,X2)|k4_xboole_0(k2_tarski(X3,X1),X2)!=k2_tarski(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_50, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_51, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_52, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_53, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_54, plain, ![X73, X74]:k6_enumset1(X73,X73,X73,X73,X73,X73,X73,X74)=k2_tarski(X73,X74), inference(variable_rename,[status(thm)],[t95_enumset1])).
cnf(c_0_55, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|k3_zfmisc_1(X1,X2,X3)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_56, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_57, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0))|r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_46]), c_0_48])).
cnf(c_0_58, plain, (k5_xboole_0(k1_enumset1(X3,X3,X1),k3_xboole_0(k1_enumset1(X3,X3,X1),X2))!=k1_enumset1(X3,X3,X1)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_50]), c_0_50]), c_0_51])).
cnf(c_0_59, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k2_xboole_0(X1,X2)))=k1_xboole_0), inference(rw,[status(thm)],[c_0_52, c_0_51])).
fof(c_0_60, plain, ![X75]:k2_xboole_0(X75,k1_xboole_0)=X75, inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_61, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X4,X2)), inference(rw,[status(thm)],[c_0_53, c_0_50])).
cnf(c_0_62, plain, (k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_63, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_55, c_0_27])).
cnf(c_0_64, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),X1)=k1_xboole_0|r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_65, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_66, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_67, plain, (k1_enumset1(X1,X1,X2)!=k1_xboole_0|~r2_hidden(X2,k2_xboole_0(k1_enumset1(X1,X1,X2),X3))), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_68, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_69, plain, (r2_hidden(X1,k1_enumset1(X2,X2,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_61])])).
fof(c_0_70, plain, ![X34, X35, X36]:((r2_hidden(k1_mcart_1(X34),X35)|~r2_hidden(X34,k2_zfmisc_1(X35,X36)))&(r2_hidden(k2_mcart_1(X34),X36)|~r2_hidden(X34,k2_zfmisc_1(X35,X36)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_mcart_1])])])).
cnf(c_0_71, plain, (k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2)=k1_enumset1(X1,X1,X2)), inference(rw,[status(thm)],[c_0_62, c_0_50])).
cnf(c_0_72, negated_conjecture, (X1=k1_xboole_0|r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_65]), c_0_66])).
cnf(c_0_73, plain, (k1_enumset1(X1,X1,X2)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_69])])).
fof(c_0_74, plain, ![X76, X77, X78]:k3_mcart_1(X76,X77,X78)=k4_tarski(k4_tarski(X76,X77),X78), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_75, plain, ![X79, X80]:(~v1_relat_1(X80)|(~r2_hidden(X79,X80)|X79=k4_tarski(k1_mcart_1(X79),k2_mcart_1(X79)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_mcart_1])])).
cnf(c_0_76, plain, (r2_hidden(k1_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_77, negated_conjecture, (r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73])).
fof(c_0_78, plain, ![X81, X82]:v1_relat_1(k2_zfmisc_1(X81,X82)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_79, negated_conjecture, (esk4_0=X3|~m1_subset_1(X1,esk1_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X3,esk3_0)|esk5_0!=k3_mcart_1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_80, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_81, plain, (X2=k4_tarski(k1_mcart_1(X2),k2_mcart_1(X2))|~v1_relat_1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_82, negated_conjecture, (r2_hidden(k1_mcart_1(esk5_0),k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_83, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_84, negated_conjecture, (esk4_0=X3|esk5_0!=k4_tarski(k4_tarski(X1,X2),X3)|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk1_0)), inference(rw,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_85, negated_conjecture, (k4_tarski(k1_mcart_1(k1_mcart_1(esk5_0)),k2_mcart_1(k1_mcart_1(esk5_0)))=k1_mcart_1(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_83])])).
cnf(c_0_86, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_87, plain, (r2_hidden(k2_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_70])).
fof(c_0_88, plain, ![X83, X84, X85, X86]:((k5_mcart_1(X83,X84,X85,X86)=k1_mcart_1(k1_mcart_1(X86))|~m1_subset_1(X86,k3_zfmisc_1(X83,X84,X85))|X85=k1_xboole_0|X84=k1_xboole_0|X83=k1_xboole_0)&((k6_mcart_1(X83,X84,X85,X86)=k2_mcart_1(k1_mcart_1(X86))|~m1_subset_1(X86,k3_zfmisc_1(X83,X84,X85))|X85=k1_xboole_0|X84=k1_xboole_0|X83=k1_xboole_0)&(k7_mcart_1(X83,X84,X85,X86)=k2_mcart_1(X86)|~m1_subset_1(X86,k3_zfmisc_1(X83,X84,X85))|X85=k1_xboole_0|X84=k1_xboole_0|X83=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_mcart_1])])])])])).
cnf(c_0_89, negated_conjecture, (esk4_0=X1|k4_tarski(k1_mcart_1(esk5_0),X1)!=esk5_0|~m1_subset_1(k2_mcart_1(k1_mcart_1(esk5_0)),esk2_0)|~m1_subset_1(k1_mcart_1(k1_mcart_1(esk5_0)),esk1_0)|~m1_subset_1(X1,esk3_0)), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_90, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_86, c_0_28])).
cnf(c_0_91, negated_conjecture, (r2_hidden(k2_mcart_1(k1_mcart_1(esk5_0)),esk2_0)), inference(spm,[status(thm)],[c_0_87, c_0_82])).
cnf(c_0_92, plain, (k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_93, negated_conjecture, (esk4_0=X1|k4_tarski(k1_mcart_1(esk5_0),X1)!=esk5_0|~m1_subset_1(k1_mcart_1(k1_mcart_1(esk5_0)),esk1_0)|~m1_subset_1(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_91])])).
cnf(c_0_94, negated_conjecture, (r2_hidden(k1_mcart_1(k1_mcart_1(esk5_0)),esk1_0)), inference(spm,[status(thm)],[c_0_76, c_0_82])).
cnf(c_0_95, negated_conjecture, (esk4_0!=k7_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_96, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_92, c_0_27])).
cnf(c_0_97, negated_conjecture, (esk4_0=X1|k4_tarski(k1_mcart_1(esk5_0),X1)!=esk5_0|~m1_subset_1(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_90]), c_0_94])])).
cnf(c_0_98, negated_conjecture, (k4_tarski(k1_mcart_1(esk5_0),k2_mcart_1(esk5_0))=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_77]), c_0_83])])).
cnf(c_0_99, negated_conjecture, (k2_mcart_1(esk5_0)!=esk4_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_31])]), c_0_48]), c_0_65]), c_0_66])).
cnf(c_0_100, negated_conjecture, (~m1_subset_1(k2_mcart_1(esk5_0),esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_99])).
cnf(c_0_101, negated_conjecture, (r2_hidden(k2_mcart_1(esk5_0),esk3_0)), inference(spm,[status(thm)],[c_0_87, c_0_77])).
cnf(c_0_102, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_90]), c_0_101])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 103
# Proof object clause steps            : 64
# Proof object formula steps           : 39
# Proof object conjectures             : 30
# Proof object clause conjectures      : 27
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 24
# Proof object simplifying inferences  : 38
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 50
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 46
# Processed clauses                    : 808
# ...of these trivial                  : 9
# ...subsumed                          : 548
# ...remaining for further processing  : 251
# Other redundant clauses eliminated   : 81
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 41
# Backward-rewritten                   : 16
# Generated clauses                    : 3280
# ...of the previous two non-trivial   : 2934
# Contextual simplify-reflections      : 2
# Paramodulations                      : 3174
# Factorizations                       : 28
# Equation resolutions                 : 81
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 138
#    Positive orientable unit clauses  : 22
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 14
#    Non-unit-clauses                  : 102
# Current number of unprocessed clauses: 1423
# ...number of literals in the above   : 5626
# Current number of archived formulas  : 0
# Current number of archived clauses   : 107
# Clause-clause subsumption calls (NU) : 2086
# Rec. Clause-clause subsumption calls : 1451
# Non-unit clause-clause subsumptions  : 214
# Unit Clause-clause subsumption calls : 208
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 18
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 33466

# -------------------------------------------------
# User time                : 0.032 s
# System time              : 0.004 s
# Total time               : 0.036 s
# Maximum resident set size: 28752 pages
