# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/5de3b349a35e4b15848b16ccee210352.E_input', t3_xboole_0)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/5de3b349a35e4b15848b16ccee210352.E_input', d7_xboole_0)).
fof(t96_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),k5_xboole_0(X1,X2)), file('data/MPTP/proofs/5de3b349a35e4b15848b16ccee210352.E_input', t96_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/5de3b349a35e4b15848b16ccee210352.E_input', t100_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/5de3b349a35e4b15848b16ccee210352.E_input', t3_boole)).
fof(t123_zfmisc_1, axiom, ![X1, X2, X3, X4]:k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)), file('data/MPTP/proofs/5de3b349a35e4b15848b16ccee210352.E_input', t123_zfmisc_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/5de3b349a35e4b15848b16ccee210352.E_input', commutativity_k3_xboole_0)).
fof(t125_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k4_xboole_0(X1,X2),X3)=k4_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&k2_zfmisc_1(X3,k4_xboole_0(X1,X2))=k4_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))), file('data/MPTP/proofs/5de3b349a35e4b15848b16ccee210352.E_input', t125_zfmisc_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/5de3b349a35e4b15848b16ccee210352.E_input', t28_xboole_1)).
fof(l97_xboole_1, axiom, ![X1, X2]:r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2)), file('data/MPTP/proofs/5de3b349a35e4b15848b16ccee210352.E_input', l97_xboole_1)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/5de3b349a35e4b15848b16ccee210352.E_input', d2_tarski)).
fof(t88_enumset1, axiom, ![X1, X2]:k4_enumset1(X1,X1,X1,X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/5de3b349a35e4b15848b16ccee210352.E_input', t88_enumset1)).
fof(t131_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(~(X1=X2)=>(r1_xboole_0(k2_zfmisc_1(k1_tarski(X1),X3),k2_zfmisc_1(k1_tarski(X2),X4))&r1_xboole_0(k2_zfmisc_1(X3,k1_tarski(X1)),k2_zfmisc_1(X4,k1_tarski(X2))))), file('data/MPTP/proofs/5de3b349a35e4b15848b16ccee210352.E_input', t131_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/5de3b349a35e4b15848b16ccee210352.E_input', t69_enumset1)).
fof(t57_zfmisc_1, axiom, ![X1, X2, X3]:~((~(r2_hidden(X1,X2))&(~(r2_hidden(X3,X2))&~(r1_xboole_0(k2_tarski(X1,X3),X2))))), file('data/MPTP/proofs/5de3b349a35e4b15848b16ccee210352.E_input', t57_zfmisc_1)).
fof(c_0_15, plain, ![X5, X6, X8, X9, X10]:(((r2_hidden(esk1_2(X5,X6),X5)|r1_xboole_0(X5,X6))&(r2_hidden(esk1_2(X5,X6),X6)|r1_xboole_0(X5,X6)))&(~r2_hidden(X10,X8)|~r2_hidden(X10,X9)|~r1_xboole_0(X8,X9))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_16, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_17, plain, (r2_hidden(esk1_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_18, plain, (r1_xboole_0(X1,X2)|~r2_hidden(esk1_2(X1,X2),X3)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_19, plain, (r2_hidden(esk1_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_20, plain, ![X11, X12]:((~r1_xboole_0(X11,X12)|k3_xboole_0(X11,X12)=k1_xboole_0)&(k3_xboole_0(X11,X12)!=k1_xboole_0|r1_xboole_0(X11,X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
fof(c_0_21, plain, ![X13, X14]:r1_tarski(k4_xboole_0(X13,X14),k5_xboole_0(X13,X14)), inference(variable_rename,[status(thm)],[t96_xboole_1])).
fof(c_0_22, plain, ![X15, X16]:k4_xboole_0(X15,X16)=k5_xboole_0(X15,k3_xboole_0(X15,X16)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_23, plain, ![X17]:k4_xboole_0(X17,k1_xboole_0)=X17, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_24, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_26, plain, ![X18, X19, X20, X21]:k2_zfmisc_1(k3_xboole_0(X18,X19),k3_xboole_0(X20,X21))=k3_xboole_0(k2_zfmisc_1(X18,X20),k2_zfmisc_1(X19,X21)), inference(variable_rename,[status(thm)],[t123_zfmisc_1])).
fof(c_0_27, plain, ![X22, X23]:k3_xboole_0(X22,X23)=k3_xboole_0(X23,X22), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_28, plain, (r1_tarski(k4_xboole_0(X1,X2),k5_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_31, plain, ![X24, X25, X26]:(k2_zfmisc_1(k4_xboole_0(X24,X25),X26)=k4_xboole_0(k2_zfmisc_1(X24,X26),k2_zfmisc_1(X25,X26))&k2_zfmisc_1(X26,k4_xboole_0(X24,X25))=k4_xboole_0(k2_zfmisc_1(X26,X24),k2_zfmisc_1(X26,X25))), inference(variable_rename,[status(thm)],[t125_zfmisc_1])).
cnf(c_0_32, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_33, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X2,X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_34, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_36, plain, ![X27, X28]:(~r1_tarski(X27,X28)|k3_xboole_0(X27,X28)=X27), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_37, plain, (r1_tarski(k5_xboole_0(X1,k3_xboole_0(X1,X2)),k5_xboole_0(X1,X2))), inference(rw,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_38, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[c_0_30, c_0_29])).
cnf(c_0_39, plain, (k2_zfmisc_1(X1,k4_xboole_0(X2,X3))=k4_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_40, plain, (k2_zfmisc_1(k4_xboole_0(X1,X2),X3)=k4_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|k3_xboole_0(X2,X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_42, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X4),k2_zfmisc_1(X3,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_34])).
cnf(c_0_43, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44, plain, (r1_tarski(X1,k5_xboole_0(X1,k1_xboole_0))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_45, plain, (k2_zfmisc_1(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))=k5_xboole_0(k2_zfmisc_1(X1,X2),k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_29]), c_0_29])).
cnf(c_0_46, plain, (k2_zfmisc_1(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X3)=k5_xboole_0(k2_zfmisc_1(X1,X3),k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_29]), c_0_29])).
cnf(c_0_47, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=k1_xboole_0|k3_xboole_0(k2_zfmisc_1(X3,X2),k2_zfmisc_1(X1,X4))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_48, plain, (k3_xboole_0(X1,k5_xboole_0(X1,k1_xboole_0))=X1), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_49, plain, (k2_zfmisc_1(k5_xboole_0(X1,k3_xboole_0(X1,X1)),X2)=k2_zfmisc_1(X1,k5_xboole_0(X2,k3_xboole_0(X2,X2)))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_50, plain, (k5_xboole_0(X1,k3_xboole_0(k1_xboole_0,X1))=X1), inference(spm,[status(thm)],[c_0_38, c_0_35])).
fof(c_0_51, plain, ![X29, X30]:r1_xboole_0(k3_xboole_0(X29,X30),k5_xboole_0(X29,X30)), inference(variable_rename,[status(thm)],[l97_xboole_1])).
cnf(c_0_52, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=k1_xboole_0|k3_xboole_0(k2_zfmisc_1(X1,X4),k2_zfmisc_1(X3,X2))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_47, c_0_35])).
cnf(c_0_53, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(k5_xboole_0(X1,k1_xboole_0),X3))=k2_zfmisc_1(X1,k3_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_34, c_0_48])).
cnf(c_0_54, plain, (k2_zfmisc_1(k1_xboole_0,k5_xboole_0(X1,k3_xboole_0(X1,X1)))=k2_zfmisc_1(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_55, plain, (r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_56, plain, (k2_zfmisc_1(X1,k3_xboole_0(X2,X3))=k1_xboole_0|k2_zfmisc_1(X1,k3_xboole_0(X3,X2))!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_53])).
cnf(c_0_57, plain, (k2_zfmisc_1(k1_xboole_0,k3_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X1)),X2))=k2_zfmisc_1(k1_xboole_0,k3_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_53])).
cnf(c_0_58, plain, (r1_xboole_0(X1,k5_xboole_0(X1,k5_xboole_0(X1,k1_xboole_0)))), inference(spm,[status(thm)],[c_0_55, c_0_48])).
cnf(c_0_59, plain, (r1_xboole_0(k3_xboole_0(X1,k3_xboole_0(X1,k1_xboole_0)),X1)), inference(spm,[status(thm)],[c_0_55, c_0_38])).
cnf(c_0_60, plain, (k2_zfmisc_1(k1_xboole_0,k3_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X2))))=k1_xboole_0|k2_zfmisc_1(k1_xboole_0,k3_xboole_0(X2,X1))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_61, plain, (k3_xboole_0(k5_xboole_0(X1,X2),k5_xboole_0(X1,k3_xboole_0(X1,X2)))=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_37]), c_0_35])).
cnf(c_0_62, plain, (k3_xboole_0(X1,k5_xboole_0(X1,k5_xboole_0(X1,k1_xboole_0)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_32, c_0_58])).
cnf(c_0_63, plain, (k5_xboole_0(k5_xboole_0(k1_xboole_0,k1_xboole_0),k1_xboole_0)=k5_xboole_0(k1_xboole_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_50, c_0_48])).
cnf(c_0_64, plain, (r1_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_24, c_0_55])).
cnf(c_0_65, plain, (k3_xboole_0(X1,k3_xboole_0(X1,k3_xboole_0(X1,k1_xboole_0)))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_59]), c_0_35])).
cnf(c_0_66, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0|k2_zfmisc_1(k1_xboole_0,k3_xboole_0(X1,k5_xboole_0(X1,X1)))!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_54])).
cnf(c_0_67, plain, (k3_xboole_0(k5_xboole_0(k1_xboole_0,k1_xboole_0),k5_xboole_0(k5_xboole_0(k1_xboole_0,k1_xboole_0),k5_xboole_0(k1_xboole_0,k1_xboole_0)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_68, plain, (k3_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_64]), c_0_35])).
cnf(c_0_69, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,k5_xboole_0(X2,k1_xboole_0)))=k2_zfmisc_1(k3_xboole_0(X1,X3),X2)), inference(spm,[status(thm)],[c_0_34, c_0_48])).
cnf(c_0_70, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,k3_xboole_0(X2,k3_xboole_0(X2,k1_xboole_0))))=k2_zfmisc_1(k3_xboole_0(X1,X3),k1_xboole_0)), inference(spm,[status(thm)],[c_0_34, c_0_65])).
cnf(c_0_71, plain, (k2_zfmisc_1(k1_xboole_0,k5_xboole_0(k1_xboole_0,k1_xboole_0))=k1_xboole_0|k2_zfmisc_1(k1_xboole_0,k1_xboole_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_72, plain, (k3_xboole_0(k1_xboole_0,k5_xboole_0(X1,k3_xboole_0(X1,k3_xboole_0(X1,k1_xboole_0))))=k1_xboole_0), inference(spm,[status(thm)],[c_0_68, c_0_65])).
cnf(c_0_73, plain, (k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,k3_xboole_0(k1_xboole_0,X1))))=k5_xboole_0(X1,k3_xboole_0(X1,k3_xboole_0(k1_xboole_0,X1)))), inference(spm,[status(thm)],[c_0_61, c_0_50])).
cnf(c_0_74, plain, (r1_xboole_0(k3_xboole_0(X1,k3_xboole_0(k1_xboole_0,X1)),X1)), inference(spm,[status(thm)],[c_0_55, c_0_50])).
cnf(c_0_75, plain, (k5_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(k3_xboole_0(X1,X1),X2))=k2_zfmisc_1(X1,k5_xboole_0(X2,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_69]), c_0_48])).
cnf(c_0_76, plain, (k2_zfmisc_1(k3_xboole_0(k1_xboole_0,X1),k1_xboole_0)=k3_xboole_0(k1_xboole_0,k2_zfmisc_1(X1,k1_xboole_0))|k2_zfmisc_1(k1_xboole_0,k1_xboole_0)!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_35]), c_0_48]), c_0_35]), c_0_48])).
cnf(c_0_77, plain, (r1_xboole_0(X1,X2)|~r2_hidden(esk1_2(X1,X2),X3)|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_16, c_0_19])).
cnf(c_0_78, plain, (k5_xboole_0(k1_xboole_0,k3_xboole_0(k1_xboole_0,k3_xboole_0(k1_xboole_0,k1_xboole_0)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_79, plain, (k3_xboole_0(X1,k3_xboole_0(X1,k3_xboole_0(k1_xboole_0,X1)))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_74]), c_0_35])).
cnf(c_0_80, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_18, c_0_17])).
cnf(c_0_81, plain, (k2_zfmisc_1(k1_xboole_0,k5_xboole_0(k1_xboole_0,k1_xboole_0))=k2_zfmisc_1(k1_xboole_0,k1_xboole_0)|k2_zfmisc_1(k1_xboole_0,k1_xboole_0)!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_50])).
fof(c_0_82, plain, ![X31, X32, X33, X34, X35, X36, X37, X38]:(((~r2_hidden(X34,X33)|(X34=X31|X34=X32)|X33!=k2_tarski(X31,X32))&((X35!=X31|r2_hidden(X35,X33)|X33!=k2_tarski(X31,X32))&(X35!=X32|r2_hidden(X35,X33)|X33!=k2_tarski(X31,X32))))&(((esk2_3(X36,X37,X38)!=X36|~r2_hidden(esk2_3(X36,X37,X38),X38)|X38=k2_tarski(X36,X37))&(esk2_3(X36,X37,X38)!=X37|~r2_hidden(esk2_3(X36,X37,X38),X38)|X38=k2_tarski(X36,X37)))&(r2_hidden(esk2_3(X36,X37,X38),X38)|(esk2_3(X36,X37,X38)=X36|esk2_3(X36,X37,X38)=X37)|X38=k2_tarski(X36,X37)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_83, plain, ![X40, X41]:k4_enumset1(X40,X40,X40,X40,X40,X41)=k2_tarski(X40,X41), inference(variable_rename,[status(thm)],[t88_enumset1])).
cnf(c_0_84, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_77, c_0_19])).
cnf(c_0_85, plain, (r1_xboole_0(k1_xboole_0,k1_xboole_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_78]), c_0_79])).
cnf(c_0_86, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_80, c_0_25])).
cnf(c_0_87, plain, (k3_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0),k2_zfmisc_1(X1,k1_xboole_0))=k2_zfmisc_1(k3_xboole_0(k1_xboole_0,X1),k1_xboole_0)|k2_zfmisc_1(k1_xboole_0,k1_xboole_0)!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_81]), c_0_35]), c_0_48]), c_0_35]), c_0_48])).
cnf(c_0_88, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_89, plain, (k4_enumset1(X1,X1,X1,X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_83])).
fof(c_0_90, negated_conjecture, ~(![X1, X2, X3, X4]:(~(X1=X2)=>(r1_xboole_0(k2_zfmisc_1(k1_tarski(X1),X3),k2_zfmisc_1(k1_tarski(X2),X4))&r1_xboole_0(k2_zfmisc_1(X3,k1_tarski(X1)),k2_zfmisc_1(X4,k1_tarski(X2)))))), inference(assume_negation,[status(cth)],[t131_zfmisc_1])).
cnf(c_0_91, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_92, plain, (r1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0),X1)|k2_zfmisc_1(k3_xboole_0(k1_xboole_0,k1_xboole_0),k1_xboole_0)!=k1_xboole_0|k2_zfmisc_1(k1_xboole_0,k1_xboole_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_93, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k4_enumset1(X2,X2,X2,X2,X2,X4)), inference(rw,[status(thm)],[c_0_88, c_0_89])).
fof(c_0_94, negated_conjecture, (esk3_0!=esk4_0&(~r1_xboole_0(k2_zfmisc_1(k1_tarski(esk3_0),esk5_0),k2_zfmisc_1(k1_tarski(esk4_0),esk6_0))|~r1_xboole_0(k2_zfmisc_1(esk5_0,k1_tarski(esk3_0)),k2_zfmisc_1(esk6_0,k1_tarski(esk4_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_90])])])])).
fof(c_0_95, plain, ![X46]:k2_tarski(X46,X46)=k1_tarski(X46), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_96, plain, (k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,k3_xboole_0(X1,k1_xboole_0))))=k5_xboole_0(X1,k3_xboole_0(X1,k3_xboole_0(X1,k1_xboole_0)))), inference(spm,[status(thm)],[c_0_61, c_0_38])).
cnf(c_0_97, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_32, c_0_91])).
cnf(c_0_98, plain, (r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_55, c_0_35])).
cnf(c_0_99, plain, (k3_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0),X1)=k1_xboole_0|k2_zfmisc_1(k3_xboole_0(k1_xboole_0,k1_xboole_0),k1_xboole_0)!=k1_xboole_0|k2_zfmisc_1(k1_xboole_0,k1_xboole_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_32, c_0_92])).
cnf(c_0_100, plain, (k2_zfmisc_1(k3_xboole_0(X1,k1_xboole_0),k1_xboole_0)=k3_xboole_0(k1_xboole_0,k2_zfmisc_1(X1,k1_xboole_0))|k2_zfmisc_1(k1_xboole_0,k1_xboole_0)!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_71]), c_0_35])).
cnf(c_0_101, plain, (r1_tarski(k2_zfmisc_1(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3))),k5_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))), inference(spm,[status(thm)],[c_0_37, c_0_45])).
cnf(c_0_102, plain, (r2_hidden(X1,k4_enumset1(X1,X1,X1,X1,X1,X2))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_93])])).
cnf(c_0_103, negated_conjecture, (~r1_xboole_0(k2_zfmisc_1(k1_tarski(esk3_0),esk5_0),k2_zfmisc_1(k1_tarski(esk4_0),esk6_0))|~r1_xboole_0(k2_zfmisc_1(esk5_0,k1_tarski(esk3_0)),k2_zfmisc_1(esk6_0,k1_tarski(esk4_0)))), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_104, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
fof(c_0_105, plain, ![X47, X48, X49]:(r2_hidden(X47,X48)|(r2_hidden(X49,X48)|r1_xboole_0(k2_tarski(X47,X49),X48))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t57_zfmisc_1])])])).
cnf(c_0_106, plain, (r1_xboole_0(k3_xboole_0(k2_zfmisc_1(X1,X2),k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))),k2_zfmisc_1(k5_xboole_0(X1,k3_xboole_0(X1,X3)),X2))), inference(spm,[status(thm)],[c_0_55, c_0_46])).
cnf(c_0_107, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_96, c_0_97]), c_0_97]), c_0_48]), c_0_97]), c_0_97])).
cnf(c_0_108, plain, (r1_xboole_0(X1,k5_xboole_0(k5_xboole_0(X1,k1_xboole_0),X1))), inference(spm,[status(thm)],[c_0_98, c_0_48])).
cnf(c_0_109, plain, (k5_xboole_0(k2_zfmisc_1(X1,X2),k3_xboole_0(k2_zfmisc_1(X3,X2),k2_zfmisc_1(X1,X2)))=k2_zfmisc_1(k5_xboole_0(X1,k3_xboole_0(X1,X3)),X2)), inference(spm,[status(thm)],[c_0_46, c_0_35])).
cnf(c_0_110, plain, (k3_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0),X1)=k1_xboole_0|k3_xboole_0(k1_xboole_0,k2_zfmisc_1(k1_xboole_0,k1_xboole_0))!=k1_xboole_0|k2_zfmisc_1(k1_xboole_0,k1_xboole_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_111, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k5_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,k1_xboole_0)))), inference(spm,[status(thm)],[c_0_101, c_0_38])).
cnf(c_0_112, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(k4_enumset1(esk1_2(X1,X2),esk1_2(X1,X2),esk1_2(X1,X2),esk1_2(X1,X2),esk1_2(X1,X2),X3),X1)), inference(spm,[status(thm)],[c_0_18, c_0_102])).
cnf(c_0_113, negated_conjecture, (~r1_xboole_0(k2_zfmisc_1(esk5_0,k4_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)),k2_zfmisc_1(esk6_0,k4_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)))|~r1_xboole_0(k2_zfmisc_1(k4_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0),esk5_0),k2_zfmisc_1(k4_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0),esk6_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_104]), c_0_104]), c_0_104]), c_0_104]), c_0_89]), c_0_89]), c_0_89]), c_0_89])).
cnf(c_0_114, plain, (r2_hidden(X1,X2)|r2_hidden(X3,X2)|r1_xboole_0(k2_tarski(X1,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_115, plain, (r1_xboole_0(k3_xboole_0(k2_zfmisc_1(X1,X2),k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(k1_xboole_0,X2))),k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_106, c_0_38])).
cnf(c_0_116, plain, (k3_xboole_0(X1,X1)=X1), inference(rw,[status(thm)],[c_0_48, c_0_107])).
cnf(c_0_117, plain, (k3_xboole_0(X1,k5_xboole_0(k5_xboole_0(X1,k1_xboole_0),X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_32, c_0_108])).
cnf(c_0_118, plain, (k5_xboole_0(k2_zfmisc_1(X1,k1_xboole_0),k1_xboole_0)=k2_zfmisc_1(X1,k1_xboole_0)|k3_xboole_0(k1_xboole_0,k2_zfmisc_1(k1_xboole_0,k1_xboole_0))!=k1_xboole_0|k2_zfmisc_1(k1_xboole_0,k1_xboole_0)!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_38])).
cnf(c_0_119, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k5_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,k1_xboole_0)))=k2_zfmisc_1(X1,X2)), inference(spm,[status(thm)],[c_0_43, c_0_111])).
cnf(c_0_120, plain, (r1_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_112, c_0_91])).
cnf(c_0_121, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(k4_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0),esk6_0),k2_zfmisc_1(k4_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0),esk5_0))!=k1_xboole_0|~r1_xboole_0(k2_zfmisc_1(esk5_0,k4_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)),k2_zfmisc_1(esk6_0,k4_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_25]), c_0_35])).
cnf(c_0_122, plain, (r2_hidden(X3,X2)|r2_hidden(X1,X2)|r1_xboole_0(k4_enumset1(X1,X1,X1,X1,X1,X3),X2)), inference(rw,[status(thm)],[c_0_114, c_0_89])).
cnf(c_0_123, plain, (r1_xboole_0(k2_zfmisc_1(k1_xboole_0,X1),k2_zfmisc_1(k1_xboole_0,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_116])).
cnf(c_0_124, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0|k3_xboole_0(k1_xboole_0,k2_zfmisc_1(k1_xboole_0,k1_xboole_0))!=k1_xboole_0|k2_zfmisc_1(k1_xboole_0,k1_xboole_0)!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_119])).
cnf(c_0_125, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_32, c_0_120])).
cnf(c_0_126, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(k4_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0),esk6_0),k2_zfmisc_1(k4_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0),esk5_0))!=k1_xboole_0|k3_xboole_0(k2_zfmisc_1(esk5_0,k4_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)),k2_zfmisc_1(esk6_0,k4_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_121, c_0_25])).
cnf(c_0_127, plain, (k3_xboole_0(k4_enumset1(X1,X1,X1,X1,X1,X2),X3)=k1_xboole_0|r2_hidden(X1,X3)|r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_32, c_0_122])).
cnf(c_0_128, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_123]), c_0_116])).
cnf(c_0_129, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0|k2_zfmisc_1(k1_xboole_0,k1_xboole_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_124, c_0_125])])).
cnf(c_0_130, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_131, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(k4_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0),esk5_0),k2_zfmisc_1(k4_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0),esk6_0))!=k1_xboole_0|k3_xboole_0(k2_zfmisc_1(esk5_0,k4_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)),k2_zfmisc_1(esk6_0,k4_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_126, c_0_42]), c_0_42])).
cnf(c_0_132, plain, (k3_xboole_0(k2_zfmisc_1(k4_enumset1(X1,X1,X1,X1,X1,X2),X3),k2_zfmisc_1(X4,X5))=k1_xboole_0|r2_hidden(X2,X4)|r2_hidden(X1,X4)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_127]), c_0_128])).
cnf(c_0_133, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_129, c_0_128])])).
cnf(c_0_134, plain, (X1=X4|X1=X3|X2!=k4_enumset1(X3,X3,X3,X3,X3,X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_130, c_0_89])).
cnf(c_0_135, negated_conjecture, (r2_hidden(esk4_0,k4_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))|k3_xboole_0(k2_zfmisc_1(esk5_0,k4_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)),k2_zfmisc_1(esk6_0,k4_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_131, c_0_132])).
cnf(c_0_136, plain, (k3_xboole_0(k2_zfmisc_1(X1,k4_enumset1(X2,X2,X2,X2,X2,X3)),k2_zfmisc_1(X4,X5))=k1_xboole_0|r2_hidden(X3,X5)|r2_hidden(X2,X5)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_127]), c_0_133])).
cnf(c_0_137, plain, (X1=X2|X1=X3|~r2_hidden(X1,k4_enumset1(X3,X3,X3,X3,X3,X2))), inference(er,[status(thm)],[c_0_134])).
cnf(c_0_138, negated_conjecture, (r2_hidden(esk4_0,k4_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_139, negated_conjecture, (esk3_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_140, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_138]), c_0_139]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 141
# Proof object clause steps            : 110
# Proof object formula steps           : 31
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 74
# Proof object simplifying inferences  : 60
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 22
# Processed clauses                    : 8081
# ...of these trivial                  : 1604
# ...subsumed                          : 5191
# ...remaining for further processing  : 1286
# Other redundant clauses eliminated   : 668
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 44
# Backward-rewritten                   : 480
# Generated clauses                    : 333656
# ...of the previous two non-trivial   : 234717
# Contextual simplify-reflections      : 3
# Paramodulations                      : 332938
# Factorizations                       : 52
# Equation resolutions                 : 668
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 737
#    Positive orientable unit clauses  : 334
#    Positive unorientable unit clauses: 11
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 387
# Current number of unprocessed clauses: 220852
# ...number of literals in the above   : 841402
# Current number of archived formulas  : 0
# Current number of archived clauses   : 549
# Clause-clause subsumption calls (NU) : 100504
# Rec. Clause-clause subsumption calls : 78440
# Non-unit clause-clause subsumptions  : 4940
# Unit Clause-clause subsumption calls : 3199
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2434
# BW rewrite match successes           : 309
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8750596

# -------------------------------------------------
# User time                : 4.851 s
# System time              : 0.113 s
# Total time               : 4.964 s
# Maximum resident set size: 219864 pages
