# No SInE strategy applied
# Auto-Mode selected heuristic H_____011_C18_F1_PI_SE_SP_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t25_waybel_0, conjecture, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(X3=X4=>((v12_waybel_0(X3,X1)=>v12_waybel_0(X4,X2))&(v13_waybel_0(X3,X1)=>v13_waybel_0(X4,X2))))))))), file('data/MPTP/proofs/5e33e4a41f7a4f1986e2243dd13b085a.E_input', t25_waybel_0)).
fof(free_g1_orders_2, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=>![X3, X4]:(g1_orders_2(X1,X2)=g1_orders_2(X3,X4)=>(X1=X3&X2=X4))), file('data/MPTP/proofs/5e33e4a41f7a4f1986e2243dd13b085a.E_input', free_g1_orders_2)).
fof(dt_u1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))), file('data/MPTP/proofs/5e33e4a41f7a4f1986e2243dd13b085a.E_input', dt_u1_orders_2)).
fof(t13_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(X3=X4=>(k3_waybel_0(X1,X3)=k3_waybel_0(X2,X4)&k4_waybel_0(X1,X3)=k4_waybel_0(X2,X4)))))))), file('data/MPTP/proofs/5e33e4a41f7a4f1986e2243dd13b085a.E_input', t13_waybel_0)).
fof(t23_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v12_waybel_0(X2,X1)<=>r1_tarski(k3_waybel_0(X1,X2),X2)))), file('data/MPTP/proofs/5e33e4a41f7a4f1986e2243dd13b085a.E_input', t23_waybel_0)).
fof(d20_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v13_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r1_orders_2(X1,X3,X4))=>r2_hidden(X4,X2))))))), file('data/MPTP/proofs/5e33e4a41f7a4f1986e2243dd13b085a.E_input', d20_waybel_0)).
fof(d9_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)<=>r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1)))))), file('data/MPTP/proofs/5e33e4a41f7a4f1986e2243dd13b085a.E_input', d9_orders_2)).
fof(c_0_7, negated_conjecture, ~(![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(X3=X4=>((v12_waybel_0(X3,X1)=>v12_waybel_0(X4,X2))&(v13_waybel_0(X3,X1)=>v13_waybel_0(X4,X2)))))))))), inference(assume_negation,[status(cth)],[t25_waybel_0])).
fof(c_0_8, plain, ![X9, X10, X11, X12]:((X9=X11|g1_orders_2(X9,X10)!=g1_orders_2(X11,X12)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(X9,X9))))&(X10=X12|g1_orders_2(X9,X10)!=g1_orders_2(X11,X12)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(X9,X9))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_orders_2])])])])).
fof(c_0_9, negated_conjecture, (l1_orders_2(esk1_0)&(l1_orders_2(esk2_0)&(g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&(esk3_0=esk4_0&(((v13_waybel_0(esk3_0,esk1_0)|v12_waybel_0(esk3_0,esk1_0))&(~v13_waybel_0(esk4_0,esk2_0)|v12_waybel_0(esk3_0,esk1_0)))&((v13_waybel_0(esk3_0,esk1_0)|~v12_waybel_0(esk4_0,esk2_0))&(~v13_waybel_0(esk4_0,esk2_0)|~v12_waybel_0(esk4_0,esk2_0)))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])).
cnf(c_0_10, plain, (X1=X2|g1_orders_2(X3,X1)!=g1_orders_2(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X3)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11, negated_conjecture, (g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12, negated_conjecture, (X1=u1_orders_2(esk2_0)|g1_orders_2(X2,X1)!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))), inference(spm,[status(thm)],[c_0_10, c_0_11])).
fof(c_0_13, plain, ![X13]:(~l1_orders_2(X13)|m1_subset_1(u1_orders_2(X13),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X13),u1_struct_0(X13))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_orders_2])])).
cnf(c_0_14, negated_conjecture, (u1_orders_2(esk2_0)=u1_orders_2(esk1_0)|~m1_subset_1(u1_orders_2(esk1_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(er,[status(thm)],[c_0_12])).
cnf(c_0_15, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_16, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, negated_conjecture, (u1_orders_2(esk2_0)=u1_orders_2(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])])).
cnf(c_0_18, plain, (X1=X2|g1_orders_2(X1,X3)!=g1_orders_2(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_19, negated_conjecture, (g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))), inference(rw,[status(thm)],[c_0_11, c_0_17])).
fof(c_0_20, plain, ![X14, X15, X16, X17]:((k3_waybel_0(X14,X16)=k3_waybel_0(X15,X17)|X16!=X17|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X15)))|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X14)))|g1_orders_2(u1_struct_0(X14),u1_orders_2(X14))!=g1_orders_2(u1_struct_0(X15),u1_orders_2(X15))|~l1_orders_2(X15)|~l1_orders_2(X14))&(k4_waybel_0(X14,X16)=k4_waybel_0(X15,X17)|X16!=X17|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X15)))|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X14)))|g1_orders_2(u1_struct_0(X14),u1_orders_2(X14))!=g1_orders_2(u1_struct_0(X15),u1_orders_2(X15))|~l1_orders_2(X15)|~l1_orders_2(X14))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_waybel_0])])])])).
cnf(c_0_21, negated_conjecture, (X1=u1_struct_0(esk2_0)|g1_orders_2(X1,X2)!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_22, plain, (k3_waybel_0(X1,X2)=k3_waybel_0(X3,X4)|X2!=X4|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))|~l1_orders_2(X3)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_23, negated_conjecture, (u1_struct_0(esk2_0)=u1_struct_0(esk1_0)|~m1_subset_1(u1_orders_2(esk1_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(er,[status(thm)],[c_0_21])).
cnf(c_0_24, plain, (k3_waybel_0(X1,X2)=k3_waybel_0(X3,X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X3)|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_25, negated_conjecture, (l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_26, negated_conjecture, (u1_struct_0(esk2_0)=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_15]), c_0_16])])).
cnf(c_0_27, negated_conjecture, (k3_waybel_0(X1,X2)=k3_waybel_0(esk2_0,X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_17]), c_0_19]), c_0_25])]), c_0_26])).
fof(c_0_28, plain, ![X18, X19]:((~v12_waybel_0(X19,X18)|r1_tarski(k3_waybel_0(X18,X19),X19)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|~l1_orders_2(X18))&(~r1_tarski(k3_waybel_0(X18,X19),X19)|v12_waybel_0(X19,X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|~l1_orders_2(X18))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_waybel_0])])])])).
fof(c_0_29, plain, ![X20, X21, X22, X23]:((~v13_waybel_0(X21,X20)|(~m1_subset_1(X22,u1_struct_0(X20))|(~m1_subset_1(X23,u1_struct_0(X20))|(~r2_hidden(X22,X21)|~r1_orders_2(X20,X22,X23)|r2_hidden(X23,X21))))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_orders_2(X20))&((m1_subset_1(esk5_2(X20,X21),u1_struct_0(X20))|v13_waybel_0(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_orders_2(X20))&((m1_subset_1(esk6_2(X20,X21),u1_struct_0(X20))|v13_waybel_0(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_orders_2(X20))&(((r2_hidden(esk5_2(X20,X21),X21)|v13_waybel_0(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_orders_2(X20))&(r1_orders_2(X20,esk5_2(X20,X21),esk6_2(X20,X21))|v13_waybel_0(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_orders_2(X20)))&(~r2_hidden(esk6_2(X20,X21),X21)|v13_waybel_0(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_orders_2(X20)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d20_waybel_0])])])])])).
cnf(c_0_30, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_31, negated_conjecture, (esk3_0=esk4_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_32, negated_conjecture, (k3_waybel_0(esk2_0,X1)=k3_waybel_0(esk1_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_27]), c_0_16])])).
cnf(c_0_33, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_34, plain, (r1_tarski(k3_waybel_0(X2,X1),X1)|~v12_waybel_0(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, plain, (m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1))|v13_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(rw,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_37, plain, (v12_waybel_0(X2,X1)|~r1_tarski(k3_waybel_0(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, negated_conjecture, (k3_waybel_0(esk2_0,esk3_0)=k3_waybel_0(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_39, negated_conjecture, (r1_tarski(k3_waybel_0(esk1_0,esk3_0),esk3_0)|~v12_waybel_0(esk3_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_33]), c_0_16])])).
cnf(c_0_40, negated_conjecture, (v13_waybel_0(esk3_0,esk1_0)|v12_waybel_0(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_41, plain, ![X26, X27, X28]:((~r1_orders_2(X26,X27,X28)|r2_hidden(k4_tarski(X27,X28),u1_orders_2(X26))|~m1_subset_1(X28,u1_struct_0(X26))|~m1_subset_1(X27,u1_struct_0(X26))|~l1_orders_2(X26))&(~r2_hidden(k4_tarski(X27,X28),u1_orders_2(X26))|r1_orders_2(X26,X27,X28)|~m1_subset_1(X28,u1_struct_0(X26))|~m1_subset_1(X27,u1_struct_0(X26))|~l1_orders_2(X26))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_orders_2])])])])).
cnf(c_0_42, negated_conjecture, (v13_waybel_0(esk3_0,esk2_0)|m1_subset_1(esk6_2(esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_25])])).
cnf(c_0_43, negated_conjecture, (v13_waybel_0(esk3_0,esk1_0)|~v12_waybel_0(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_44, negated_conjecture, (v12_waybel_0(esk3_0,esk2_0)|~r1_tarski(k3_waybel_0(esk1_0,esk3_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_26]), c_0_33]), c_0_25])])).
cnf(c_0_45, negated_conjecture, (r1_tarski(k3_waybel_0(esk1_0,esk3_0),esk3_0)|v13_waybel_0(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_46, plain, (r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_47, plain, (r2_hidden(X4,X1)|~v13_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,X1)|~r1_orders_2(X2,X3,X4)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_48, negated_conjecture, (v13_waybel_0(esk3_0,esk2_0)|m1_subset_1(esk6_2(esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(rw,[status(thm)],[c_0_42, c_0_26])).
cnf(c_0_49, negated_conjecture, (v13_waybel_0(esk3_0,esk1_0)|~v12_waybel_0(esk3_0,esk2_0)), inference(rw,[status(thm)],[c_0_43, c_0_31])).
cnf(c_0_50, negated_conjecture, (v13_waybel_0(esk3_0,esk1_0)|v12_waybel_0(esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_51, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|v13_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_52, negated_conjecture, (r2_hidden(k4_tarski(X1,esk6_2(esk2_0,esk3_0)),u1_orders_2(esk2_0))|v13_waybel_0(esk3_0,esk2_0)|~r1_orders_2(esk2_0,X1,esk6_2(esk2_0,esk3_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_42]), c_0_25])])).
cnf(c_0_53, negated_conjecture, (r2_hidden(esk6_2(esk2_0,esk3_0),X1)|v13_waybel_0(esk3_0,esk2_0)|~r1_orders_2(esk1_0,X2,esk6_2(esk2_0,esk3_0))|~r2_hidden(X2,X1)|~v13_waybel_0(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_16])])).
cnf(c_0_54, negated_conjecture, (v13_waybel_0(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_55, plain, (r2_hidden(esk5_2(X1,X2),X2)|v13_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_56, plain, (r1_orders_2(X3,X1,X2)|~r2_hidden(k4_tarski(X1,X2),u1_orders_2(X3))|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_57, negated_conjecture, (v13_waybel_0(esk3_0,esk2_0)|m1_subset_1(esk5_2(esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_36]), c_0_25])])).
cnf(c_0_58, negated_conjecture, (r2_hidden(k4_tarski(X1,esk6_2(esk2_0,esk3_0)),u1_orders_2(esk1_0))|v13_waybel_0(esk3_0,esk2_0)|~r1_orders_2(esk2_0,X1,esk6_2(esk2_0,esk3_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(rw,[status(thm)],[c_0_52, c_0_17])).
cnf(c_0_59, plain, (r1_orders_2(X1,esk5_2(X1,X2),esk6_2(X1,X2))|v13_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_60, negated_conjecture, (r2_hidden(esk6_2(esk2_0,esk3_0),esk3_0)|v13_waybel_0(esk3_0,esk2_0)|~r1_orders_2(esk1_0,X1,esk6_2(esk2_0,esk3_0))|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_33])])).
cnf(c_0_61, negated_conjecture, (r2_hidden(esk5_2(esk2_0,esk3_0),esk3_0)|v13_waybel_0(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_36]), c_0_25])])).
cnf(c_0_62, negated_conjecture, (r1_orders_2(esk1_0,X1,esk6_2(esk2_0,esk3_0))|v13_waybel_0(esk3_0,esk2_0)|~r2_hidden(k4_tarski(X1,esk6_2(esk2_0,esk3_0)),u1_orders_2(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_48]), c_0_16])])).
cnf(c_0_63, negated_conjecture, (v13_waybel_0(esk3_0,esk2_0)|m1_subset_1(esk5_2(esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(rw,[status(thm)],[c_0_57, c_0_26])).
cnf(c_0_64, negated_conjecture, (r2_hidden(k4_tarski(esk5_2(esk2_0,esk3_0),esk6_2(esk2_0,esk3_0)),u1_orders_2(esk1_0))|v13_waybel_0(esk3_0,esk2_0)|~r1_orders_2(esk2_0,esk5_2(esk2_0,esk3_0),esk6_2(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_58, c_0_57])).
cnf(c_0_65, negated_conjecture, (r1_orders_2(esk2_0,esk5_2(esk2_0,esk3_0),esk6_2(esk2_0,esk3_0))|v13_waybel_0(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_36]), c_0_25])])).
cnf(c_0_66, negated_conjecture, (r2_hidden(esk6_2(esk2_0,esk3_0),esk3_0)|v13_waybel_0(esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk5_2(esk2_0,esk3_0),esk6_2(esk2_0,esk3_0))|~m1_subset_1(esk5_2(esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_67, negated_conjecture, (r1_orders_2(esk1_0,esk5_2(esk2_0,esk3_0),esk6_2(esk2_0,esk3_0))|v13_waybel_0(esk3_0,esk2_0)|~r2_hidden(k4_tarski(esk5_2(esk2_0,esk3_0),esk6_2(esk2_0,esk3_0)),u1_orders_2(esk1_0))), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_68, negated_conjecture, (r2_hidden(k4_tarski(esk5_2(esk2_0,esk3_0),esk6_2(esk2_0,esk3_0)),u1_orders_2(esk1_0))|v13_waybel_0(esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_69, negated_conjecture, (r2_hidden(esk6_2(esk2_0,esk3_0),esk3_0)|v13_waybel_0(esk3_0,esk2_0)|~r1_orders_2(esk1_0,esk5_2(esk2_0,esk3_0),esk6_2(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_66, c_0_63])).
cnf(c_0_70, negated_conjecture, (r1_orders_2(esk1_0,esk5_2(esk2_0,esk3_0),esk6_2(esk2_0,esk3_0))|v13_waybel_0(esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_71, negated_conjecture, (v12_waybel_0(esk3_0,esk1_0)|~v13_waybel_0(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_72, plain, (v13_waybel_0(X2,X1)|~r2_hidden(esk6_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_73, negated_conjecture, (r2_hidden(esk6_2(esk2_0,esk3_0),esk3_0)|v13_waybel_0(esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_74, negated_conjecture, (v12_waybel_0(esk3_0,esk1_0)|~v13_waybel_0(esk3_0,esk2_0)), inference(rw,[status(thm)],[c_0_71, c_0_31])).
cnf(c_0_75, negated_conjecture, (v13_waybel_0(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_26]), c_0_33]), c_0_25])])).
cnf(c_0_76, negated_conjecture, (~v13_waybel_0(esk4_0,esk2_0)|~v12_waybel_0(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_77, negated_conjecture, (v12_waybel_0(esk3_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_75])])).
cnf(c_0_78, negated_conjecture, (~v13_waybel_0(esk3_0,esk2_0)|~v12_waybel_0(esk3_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_31]), c_0_31])).
cnf(c_0_79, negated_conjecture, (r1_tarski(k3_waybel_0(esk1_0,esk3_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_77])])).
cnf(c_0_80, negated_conjecture, (~v12_waybel_0(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_75])])).
cnf(c_0_81, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_79])]), c_0_80]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 82
# Proof object clause steps            : 67
# Proof object formula steps           : 15
# Proof object conjectures             : 55
# Proof object clause conjectures      : 52
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 30
# Proof object simplifying inferences  : 55
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 25
# Processed clauses                    : 150
# ...of these trivial                  : 1
# ...subsumed                          : 23
# ...remaining for further processing  : 126
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 80
# Generated clauses                    : 153
# ...of the previous two non-trivial   : 171
# Contextual simplify-reflections      : 0
# Paramodulations                      : 144
# Factorizations                       : 0
# Equation resolutions                 : 9
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 38
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 25
# Current number of unprocessed clauses: 16
# ...number of literals in the above   : 68
# Current number of archived formulas  : 0
# Current number of archived clauses   : 86
# Clause-clause subsumption calls (NU) : 269
# Rec. Clause-clause subsumption calls : 206
# Non-unit clause-clause subsumptions  : 29
# Unit Clause-clause subsumption calls : 21
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 8
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6454

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.003 s
# Total time               : 0.013 s
# Maximum resident set size: 28628 pages
