# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PI_S0a
# and selection function SelectNegativeLiterals.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t50_relset_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>~((~(k2_relset_1(X2,X1,X3)=k1_xboole_0)&![X4]:(m1_subset_1(X4,X2)=>~(r2_hidden(X4,k1_relset_1(X2,X1,X3))))))))), file('data/MPTP/proofs/5e4dfeec1ea1454b8c5489d90fb802b6.E_input', t50_relset_1)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/5e4dfeec1ea1454b8c5489d90fb802b6.E_input', redefinition_k2_relset_1)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/5e4dfeec1ea1454b8c5489d90fb802b6.E_input', d4_relat_1)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/5e4dfeec1ea1454b8c5489d90fb802b6.E_input', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/5e4dfeec1ea1454b8c5489d90fb802b6.E_input', fc6_relat_1)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/5e4dfeec1ea1454b8c5489d90fb802b6.E_input', l3_subset_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/5e4dfeec1ea1454b8c5489d90fb802b6.E_input', t7_xboole_0)).
fof(t65_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>(k1_relat_1(X1)=k1_xboole_0<=>k2_relat_1(X1)=k1_xboole_0)), file('data/MPTP/proofs/5e4dfeec1ea1454b8c5489d90fb802b6.E_input', t65_relat_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/5e4dfeec1ea1454b8c5489d90fb802b6.E_input', redefinition_k1_relset_1)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/5e4dfeec1ea1454b8c5489d90fb802b6.E_input', l54_zfmisc_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/5e4dfeec1ea1454b8c5489d90fb802b6.E_input', t1_subset)).
fof(c_0_11, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>~((~(k2_relset_1(X2,X1,X3)=k1_xboole_0)&![X4]:(m1_subset_1(X4,X2)=>~(r2_hidden(X4,k1_relset_1(X2,X1,X3)))))))))), inference(assume_negation,[status(cth)],[t50_relset_1])).
fof(c_0_12, plain, ![X11, X12, X13]:(~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12)))|k2_relset_1(X11,X12,X13)=k2_relat_1(X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
fof(c_0_13, negated_conjecture, ![X8]:(~v1_xboole_0(esk1_0)&(~v1_xboole_0(esk2_0)&(m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))&(k2_relset_1(esk2_0,esk1_0,esk3_0)!=k1_xboole_0&(~m1_subset_1(X8,esk2_0)|~r2_hidden(X8,k1_relset_1(esk2_0,esk1_0,esk3_0))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])])).
fof(c_0_14, plain, ![X31, X32, X33, X35, X36, X37, X38, X40]:(((~r2_hidden(X33,X32)|r2_hidden(k4_tarski(X33,esk5_3(X31,X32,X33)),X31)|X32!=k1_relat_1(X31))&(~r2_hidden(k4_tarski(X35,X36),X31)|r2_hidden(X35,X32)|X32!=k1_relat_1(X31)))&((~r2_hidden(esk6_2(X37,X38),X38)|~r2_hidden(k4_tarski(esk6_2(X37,X38),X40),X37)|X38=k1_relat_1(X37))&(r2_hidden(esk6_2(X37,X38),X38)|r2_hidden(k4_tarski(esk6_2(X37,X38),esk7_2(X37,X38)),X37)|X38=k1_relat_1(X37)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
fof(c_0_15, plain, ![X18, X19]:(~v1_relat_1(X18)|(~m1_subset_1(X19,k1_zfmisc_1(X18))|v1_relat_1(X19))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_16, plain, ![X9, X10]:v1_relat_1(k2_zfmisc_1(X9,X10)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_17, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_19, plain, ![X20, X21, X22]:(~m1_subset_1(X21,k1_zfmisc_1(X20))|(~r2_hidden(X22,X21)|r2_hidden(X22,X20))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
cnf(c_0_20, plain, (r2_hidden(k4_tarski(X1,esk5_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_21, plain, ![X23]:(X23=k1_xboole_0|r2_hidden(esk4_1(X23),X23)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
fof(c_0_22, plain, ![X17]:((k1_relat_1(X17)!=k1_xboole_0|k2_relat_1(X17)=k1_xboole_0|~v1_relat_1(X17))&(k2_relat_1(X17)!=k1_xboole_0|k1_relat_1(X17)=k1_xboole_0|~v1_relat_1(X17))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t65_relat_1])])])).
cnf(c_0_23, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, negated_conjecture, (k2_relset_1(esk2_0,esk1_0,esk3_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_26, negated_conjecture, (k2_relset_1(esk2_0,esk1_0,esk3_0)=k2_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
fof(c_0_27, plain, ![X14, X15, X16]:(~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15)))|k1_relset_1(X14,X15,X16)=k1_relat_1(X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_28, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (r2_hidden(k4_tarski(X1,esk5_3(X2,k1_relat_1(X2),X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_30, plain, (X1=k1_xboole_0|r2_hidden(esk4_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, plain, (k2_relat_1(X1)=k1_xboole_0|k1_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32, negated_conjecture, (v1_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_18]), c_0_24])])).
cnf(c_0_33, negated_conjecture, (k2_relat_1(esk3_0)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_34, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_35, plain, ![X25, X26, X27, X28]:(((r2_hidden(X25,X27)|~r2_hidden(k4_tarski(X25,X26),k2_zfmisc_1(X27,X28)))&(r2_hidden(X26,X28)|~r2_hidden(k4_tarski(X25,X26),k2_zfmisc_1(X27,X28))))&(~r2_hidden(X25,X27)|~r2_hidden(X26,X28)|r2_hidden(k4_tarski(X25,X26),k2_zfmisc_1(X27,X28)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
cnf(c_0_36, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(esk2_0,esk1_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_28, c_0_18])).
cnf(c_0_37, plain, (k1_relat_1(X1)=k1_xboole_0|r2_hidden(k4_tarski(esk4_1(k1_relat_1(X1)),esk5_3(X1,k1_relat_1(X1),esk4_1(k1_relat_1(X1)))),X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_38, negated_conjecture, (k1_relat_1(esk3_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])).
cnf(c_0_39, negated_conjecture, (~m1_subset_1(X1,esk2_0)|~r2_hidden(X1,k1_relset_1(esk2_0,esk1_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_40, negated_conjecture, (k1_relset_1(esk2_0,esk1_0,esk3_0)=k1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_34, c_0_18])).
fof(c_0_41, plain, ![X29, X30]:(~r2_hidden(X29,X30)|m1_subset_1(X29,X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_42, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_43, negated_conjecture, (r2_hidden(k4_tarski(esk4_1(k1_relat_1(esk3_0)),esk5_3(esk3_0,k1_relat_1(esk3_0),esk4_1(k1_relat_1(esk3_0)))),k2_zfmisc_1(esk2_0,esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_44, negated_conjecture, (~r2_hidden(X1,k1_relat_1(esk3_0))|~m1_subset_1(X1,esk2_0)), inference(rw,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_45, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk4_1(k1_relat_1(esk3_0)),esk2_0)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_47, negated_conjecture, (~m1_subset_1(esk4_1(k1_relat_1(esk3_0)),esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_30]), c_0_38])).
cnf(c_0_48, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 49
# Proof object clause steps            : 26
# Proof object formula steps           : 23
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 21
# Processed clauses                    : 596
# ...of these trivial                  : 0
# ...subsumed                          : 146
# ...remaining for further processing  : 450
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 38
# Backward-rewritten                   : 4
# Generated clauses                    : 7826
# ...of the previous two non-trivial   : 7762
# Contextual simplify-reflections      : 1
# Paramodulations                      : 7783
# Factorizations                       : 6
# Equation resolutions                 : 37
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 408
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 396
# Current number of unprocessed clauses: 7168
# ...number of literals in the above   : 30582
# Current number of archived formulas  : 0
# Current number of archived clauses   : 42
# Clause-clause subsumption calls (NU) : 18087
# Rec. Clause-clause subsumption calls : 6385
# Non-unit clause-clause subsumptions  : 184
# Unit Clause-clause subsumption calls : 243
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 203835

# -------------------------------------------------
# User time                : 0.095 s
# System time              : 0.004 s
# Total time               : 0.099 s
# Maximum resident set size: 28700 pages
