# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d13_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(m1_yellow_0(X2,X1)<=>(r1_tarski(u1_struct_0(X2),u1_struct_0(X1))&r1_tarski(u1_orders_2(X2),u1_orders_2(X1)))))), file('data/MPTP/proofs/5e51cecd94bc4e9d8a65e29d98828540.E_input', d13_yellow_0)).
fof(dt_m1_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_yellow_0(X2,X1)=>l1_orders_2(X2))), file('data/MPTP/proofs/5e51cecd94bc4e9d8a65e29d98828540.E_input', dt_m1_yellow_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/5e51cecd94bc4e9d8a65e29d98828540.E_input', d3_tarski)).
fof(t16_yellow_6, conjecture, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_yellow_0(X2,X1)=>![X3]:(m1_yellow_0(X3,X2)=>m1_yellow_0(X3,X1)))), file('data/MPTP/proofs/5e51cecd94bc4e9d8a65e29d98828540.E_input', t16_yellow_6)).
fof(c_0_4, plain, ![X13, X14]:(((r1_tarski(u1_struct_0(X14),u1_struct_0(X13))|~m1_yellow_0(X14,X13)|~l1_orders_2(X14)|~l1_orders_2(X13))&(r1_tarski(u1_orders_2(X14),u1_orders_2(X13))|~m1_yellow_0(X14,X13)|~l1_orders_2(X14)|~l1_orders_2(X13)))&(~r1_tarski(u1_struct_0(X14),u1_struct_0(X13))|~r1_tarski(u1_orders_2(X14),u1_orders_2(X13))|m1_yellow_0(X14,X13)|~l1_orders_2(X14)|~l1_orders_2(X13))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d13_yellow_0])])])])).
fof(c_0_5, plain, ![X15, X16]:(~l1_orders_2(X15)|(~m1_yellow_0(X16,X15)|l1_orders_2(X16))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_yellow_0])])])).
fof(c_0_6, plain, ![X7, X8, X9, X10, X11]:((~r1_tarski(X7,X8)|(~r2_hidden(X9,X7)|r2_hidden(X9,X8)))&((r2_hidden(esk4_2(X10,X11),X10)|r1_tarski(X10,X11))&(~r2_hidden(esk4_2(X10,X11),X11)|r1_tarski(X10,X11)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_7, plain, (r1_tarski(u1_orders_2(X1),u1_orders_2(X2))|~m1_yellow_0(X1,X2)|~l1_orders_2(X1)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, plain, (l1_orders_2(X2)|~l1_orders_2(X1)|~m1_yellow_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_yellow_0(X1,X2)|~l1_orders_2(X1)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11, plain, (r1_tarski(u1_orders_2(X1),u1_orders_2(X2))|~m1_yellow_0(X1,X2)|~l1_orders_2(X2)), inference(csr,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_12, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_yellow_0(X1,X2)|~l1_orders_2(X2)), inference(csr,[status(thm)],[c_0_9, c_0_8])).
cnf(c_0_13, plain, (r2_hidden(X1,u1_orders_2(X2))|~r2_hidden(X1,u1_orders_2(X3))|~m1_yellow_0(X3,X2)|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_14, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_15, negated_conjecture, ~(![X1]:(l1_orders_2(X1)=>![X2]:(m1_yellow_0(X2,X1)=>![X3]:(m1_yellow_0(X3,X2)=>m1_yellow_0(X3,X1))))), inference(assume_negation,[status(cth)],[t16_yellow_6])).
cnf(c_0_16, plain, (r2_hidden(X1,u1_struct_0(X2))|~r2_hidden(X1,u1_struct_0(X3))|~m1_yellow_0(X3,X2)|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_10, c_0_12])).
cnf(c_0_17, plain, (r2_hidden(esk4_2(u1_orders_2(X1),X2),u1_orders_2(X3))|r1_tarski(u1_orders_2(X1),X2)|~m1_yellow_0(X1,X3)|~l1_orders_2(X3)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
fof(c_0_18, negated_conjecture, (l1_orders_2(esk1_0)&(m1_yellow_0(esk2_0,esk1_0)&(m1_yellow_0(esk3_0,esk2_0)&~m1_yellow_0(esk3_0,esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).
cnf(c_0_19, plain, (r2_hidden(esk4_2(u1_struct_0(X1),X2),u1_struct_0(X3))|r1_tarski(u1_struct_0(X1),X2)|~m1_yellow_0(X1,X3)|~l1_orders_2(X3)), inference(spm,[status(thm)],[c_0_16, c_0_14])).
cnf(c_0_20, plain, (r2_hidden(esk4_2(u1_orders_2(X1),X2),u1_orders_2(X3))|r1_tarski(u1_orders_2(X1),X2)|~m1_yellow_0(X4,X3)|~m1_yellow_0(X1,X4)|~l1_orders_2(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_17]), c_0_8])).
cnf(c_0_21, negated_conjecture, (m1_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, plain, (r2_hidden(esk4_2(u1_struct_0(X1),X2),u1_struct_0(X3))|r1_tarski(u1_struct_0(X1),X2)|~m1_yellow_0(X4,X3)|~m1_yellow_0(X1,X4)|~l1_orders_2(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_19]), c_0_8])).
cnf(c_0_24, plain, (r1_tarski(X1,X2)|~r2_hidden(esk4_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_25, negated_conjecture, (r2_hidden(esk4_2(u1_orders_2(X1),X2),u1_orders_2(esk1_0))|r1_tarski(u1_orders_2(X1),X2)|~m1_yellow_0(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])])).
cnf(c_0_26, negated_conjecture, (r2_hidden(esk4_2(u1_struct_0(X1),X2),u1_struct_0(esk1_0))|r1_tarski(u1_struct_0(X1),X2)|~m1_yellow_0(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_21]), c_0_22])])).
cnf(c_0_27, plain, (m1_yellow_0(X1,X2)|~r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~r1_tarski(u1_orders_2(X1),u1_orders_2(X2))|~l1_orders_2(X1)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_28, negated_conjecture, (r1_tarski(u1_orders_2(X1),u1_orders_2(esk1_0))|~m1_yellow_0(X1,esk2_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk1_0))|~m1_yellow_0(X1,esk2_0)), inference(spm,[status(thm)],[c_0_24, c_0_26])).
cnf(c_0_30, negated_conjecture, (m1_yellow_0(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_31, negated_conjecture, (l1_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_21]), c_0_22])])).
cnf(c_0_32, negated_conjecture, (m1_yellow_0(X1,esk1_0)|~m1_yellow_0(X1,esk2_0)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_22])]), c_0_29])).
cnf(c_0_33, negated_conjecture, (l1_orders_2(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_30]), c_0_31])])).
cnf(c_0_34, negated_conjecture, (~m1_yellow_0(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_35, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_30]), c_0_33])]), c_0_34]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 36
# Proof object clause steps            : 27
# Proof object formula steps           : 9
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 11
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 11
# Processed clauses                    : 63
# ...of these trivial                  : 0
# ...subsumed                          : 9
# ...remaining for further processing  : 54
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 74
# ...of the previous two non-trivial   : 68
# Contextual simplify-reflections      : 6
# Paramodulations                      : 74
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 43
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 36
# Current number of unprocessed clauses: 27
# ...number of literals in the above   : 119
# Current number of archived formulas  : 0
# Current number of archived clauses   : 11
# Clause-clause subsumption calls (NU) : 268
# Rec. Clause-clause subsumption calls : 118
# Non-unit clause-clause subsumptions  : 15
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2225

# -------------------------------------------------
# User time                : 0.004 s
# System time              : 0.005 s
# Total time               : 0.010 s
# Maximum resident set size: 28620 pages
