# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', t1_boole)).
fof(d2_xboole_0, axiom, k1_xboole_0=o_0_0_xboole_0, file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', d2_xboole_0)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', commutativity_k2_xboole_0)).
fof(t4_boole, axiom, ![X1]:k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0, file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', t4_boole)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', l32_xboole_1)).
fof(t40_xboole_1, axiom, ![X1, X2]:k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', t40_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', t28_xboole_1)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', t41_xboole_1)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', t17_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', commutativity_k3_xboole_0)).
fof(t16_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3)), file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', t16_xboole_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', idempotence_k3_xboole_0)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', d5_xboole_0)).
fof(t47_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', t47_xboole_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', t3_xboole_0)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', t36_xboole_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', d7_xboole_0)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', t3_boole)).
fof(t20_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&(r1_tarski(X1,X3)&![X4]:((r1_tarski(X4,X2)&r1_tarski(X4,X3))=>r1_tarski(X4,X1))))=>X1=k3_xboole_0(X2,X3)), file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', t20_xboole_1)).
fof(t48_xboole_1, conjecture, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/5e6299b530a04b8a8a64c5a7d2826139.E_input', t48_xboole_1)).
fof(c_0_20, plain, ![X5]:k2_xboole_0(X5,k1_xboole_0)=X5, inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_21, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_22, plain, (k1_xboole_0=o_0_0_xboole_0), inference(split_conjunct,[status(thm)],[d2_xboole_0])).
fof(c_0_23, plain, ![X6, X7]:k2_xboole_0(X6,X7)=k2_xboole_0(X7,X6), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_24, plain, ![X8]:k4_xboole_0(k1_xboole_0,X8)=k1_xboole_0, inference(variable_rename,[status(thm)],[t4_boole])).
fof(c_0_25, plain, ![X9, X10]:((k4_xboole_0(X9,X10)!=k1_xboole_0|r1_tarski(X9,X10))&(~r1_tarski(X9,X10)|k4_xboole_0(X9,X10)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
fof(c_0_26, plain, ![X11, X12]:k4_xboole_0(k2_xboole_0(X11,X12),X12)=k4_xboole_0(X11,X12), inference(variable_rename,[status(thm)],[t40_xboole_1])).
cnf(c_0_27, plain, (k2_xboole_0(X1,o_0_0_xboole_0)=X1), inference(rw,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_30, plain, ![X13, X14]:(~r1_tarski(X13,X14)|k3_xboole_0(X13,X14)=X13), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_31, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, plain, (k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, plain, (k2_xboole_0(o_0_0_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34, plain, (k4_xboole_0(o_0_0_xboole_0,X1)=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_22]), c_0_22])).
fof(c_0_35, plain, ![X15, X16, X17]:k4_xboole_0(k4_xboole_0(X15,X16),X17)=k4_xboole_0(X15,k2_xboole_0(X16,X17)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
fof(c_0_36, plain, ![X18, X19]:r1_tarski(k3_xboole_0(X18,X19),X18), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_37, plain, ![X20, X21]:k3_xboole_0(X20,X21)=k3_xboole_0(X21,X20), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_38, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=o_0_0_xboole_0), inference(rw,[status(thm)],[c_0_31, c_0_22])).
cnf(c_0_40, plain, (k4_xboole_0(X1,X1)=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])).
cnf(c_0_41, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_42, plain, ![X22, X23, X24]:k3_xboole_0(k3_xboole_0(X22,X23),X24)=k3_xboole_0(X22,k3_xboole_0(X23,X24)), inference(variable_rename,[status(thm)],[t16_xboole_1])).
fof(c_0_43, plain, ![X25]:k3_xboole_0(X25,X25)=X25, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
fof(c_0_44, plain, ![X26, X27, X28, X29, X30, X31, X32, X33]:((((r2_hidden(X29,X26)|~r2_hidden(X29,X28)|X28!=k4_xboole_0(X26,X27))&(~r2_hidden(X29,X27)|~r2_hidden(X29,X28)|X28!=k4_xboole_0(X26,X27)))&(~r2_hidden(X30,X26)|r2_hidden(X30,X27)|r2_hidden(X30,X28)|X28!=k4_xboole_0(X26,X27)))&((~r2_hidden(esk1_3(X31,X32,X33),X33)|(~r2_hidden(esk1_3(X31,X32,X33),X31)|r2_hidden(esk1_3(X31,X32,X33),X32))|X33=k4_xboole_0(X31,X32))&((r2_hidden(esk1_3(X31,X32,X33),X31)|r2_hidden(esk1_3(X31,X32,X33),X33)|X33=k4_xboole_0(X31,X32))&(~r2_hidden(esk1_3(X31,X32,X33),X32)|r2_hidden(esk1_3(X31,X32,X33),X33)|X33=k4_xboole_0(X31,X32))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_45, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_46, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_47, plain, (k3_xboole_0(X1,X2)=X1|k4_xboole_0(X1,X2)!=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_48, plain, (k4_xboole_0(X1,k2_xboole_0(X2,k4_xboole_0(X1,X2)))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_49, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_50, plain, ![X35, X36]:k4_xboole_0(X35,k3_xboole_0(X35,X36))=k4_xboole_0(X35,X36), inference(variable_rename,[status(thm)],[t47_xboole_1])).
cnf(c_0_51, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_52, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_53, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_54, plain, ![X37, X38, X40, X41, X42]:(((r2_hidden(esk2_2(X37,X38),X37)|r1_xboole_0(X37,X38))&(r2_hidden(esk2_2(X37,X38),X38)|r1_xboole_0(X37,X38)))&(~r2_hidden(X42,X40)|~r2_hidden(X42,X41)|~r1_xboole_0(X40,X41))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_55, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_56, plain, (k3_xboole_0(X1,k2_xboole_0(X2,k4_xboole_0(X1,X2)))=X1), inference(spm,[status(thm)],[c_0_47, c_0_48])).
fof(c_0_57, plain, ![X43, X44]:r1_tarski(k4_xboole_0(X43,X44),X43), inference(variable_rename,[status(thm)],[t36_xboole_1])).
cnf(c_0_58, plain, (k4_xboole_0(X1,X2)=o_0_0_xboole_0|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_49, c_0_22])).
cnf(c_0_59, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_60, plain, (k3_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
fof(c_0_61, plain, ![X45, X46]:((~r1_xboole_0(X45,X46)|k3_xboole_0(X45,X46)=k1_xboole_0)&(k3_xboole_0(X45,X46)!=k1_xboole_0|r1_xboole_0(X45,X46))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_62, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_53])).
cnf(c_0_63, plain, (r2_hidden(esk2_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_64, plain, (r1_tarski(X1,k2_xboole_0(X2,k4_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_65, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_66, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X2)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_58, c_0_55])).
cnf(c_0_67, plain, (k4_xboole_0(X1,k3_xboole_0(X2,X1))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_59, c_0_46])).
cnf(c_0_68, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_60, c_0_46])).
cnf(c_0_69, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_70, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r2_hidden(esk2_2(X1,k4_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_71, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
fof(c_0_72, plain, ![X47]:k4_xboole_0(X47,k1_xboole_0)=X47, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_73, plain, (r1_tarski(k4_xboole_0(X1,X2),k2_xboole_0(X3,k4_xboole_0(X1,k2_xboole_0(X2,X3))))), inference(spm,[status(thm)],[c_0_64, c_0_41])).
cnf(c_0_74, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X1)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_58, c_0_65])).
cnf(c_0_75, plain, (k4_xboole_0(k3_xboole_0(X1,k3_xboole_0(X2,X3)),X3)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_66, c_0_51])).
cnf(c_0_76, plain, (k4_xboole_0(X1,k3_xboole_0(X2,k3_xboole_0(X3,X1)))=k4_xboole_0(X1,k3_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_67, c_0_51])).
cnf(c_0_77, plain, (k3_xboole_0(X1,k3_xboole_0(X2,k3_xboole_0(X1,X3)))=k3_xboole_0(X2,k3_xboole_0(X1,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_68]), c_0_51]), c_0_51])).
cnf(c_0_78, plain, (k3_xboole_0(X1,X2)=o_0_0_xboole_0|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_69, c_0_22])).
cnf(c_0_79, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_80, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_81, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,k2_xboole_0(X2,X3)),X2),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_66]), c_0_27])).
cnf(c_0_82, plain, (k3_xboole_0(X1,k3_xboole_0(X2,k3_xboole_0(X3,X1)))=k3_xboole_0(X2,k3_xboole_0(X3,X1))), inference(spm,[status(thm)],[c_0_68, c_0_51])).
cnf(c_0_83, plain, (k4_xboole_0(X1,k2_xboole_0(X2,X1))=o_0_0_xboole_0), inference(rw,[status(thm)],[c_0_74, c_0_41])).
cnf(c_0_84, plain, (k4_xboole_0(k3_xboole_0(X1,X2),k2_xboole_0(X3,k4_xboole_0(X2,X3)))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_75, c_0_56])).
cnf(c_0_85, plain, (k4_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3))=k4_xboole_0(k3_xboole_0(X1,X2),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_67])).
cnf(c_0_86, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X1))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_87, plain, (k4_xboole_0(X1,o_0_0_xboole_0)=X1), inference(rw,[status(thm)],[c_0_80, c_0_22])).
fof(c_0_88, plain, ![X48, X49, X50]:(((r1_tarski(esk3_3(X48,X49,X50),X49)|~r1_tarski(X48,X50)|~r1_tarski(X48,X49)|X48=k3_xboole_0(X49,X50))&(r1_tarski(esk3_3(X48,X49,X50),X50)|~r1_tarski(X48,X50)|~r1_tarski(X48,X49)|X48=k3_xboole_0(X49,X50)))&(~r1_tarski(esk3_3(X48,X49,X50),X48)|~r1_tarski(X48,X50)|~r1_tarski(X48,X49)|X48=k3_xboole_0(X49,X50))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t20_xboole_1])])])])).
cnf(c_0_89, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,k2_xboole_0(X2,X3)),X3),X2)), inference(spm,[status(thm)],[c_0_81, c_0_28])).
cnf(c_0_90, plain, (k4_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X2,X3))=k4_xboole_0(k3_xboole_0(X1,X2),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_82]), c_0_67])).
cnf(c_0_91, plain, (k3_xboole_0(X1,k2_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_47, c_0_83])).
cnf(c_0_92, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,X2),X3),k4_xboole_0(X2,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_84]), c_0_27])).
cnf(c_0_93, plain, (k4_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X3,X1))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_87])).
fof(c_0_94, negated_conjecture, ~(![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(assume_negation,[status(cth)],[t48_xboole_1])).
cnf(c_0_95, plain, (X1=k3_xboole_0(X2,X3)|~r1_tarski(esk3_3(X1,X2,X3),X1)|~r1_tarski(X1,X3)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_96, plain, (r1_tarski(esk3_3(X1,X2,X3),X3)|X1=k3_xboole_0(X2,X3)|~r1_tarski(X1,X3)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_97, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_65, c_0_87])).
cnf(c_0_98, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_89, c_0_56])).
cnf(c_0_99, plain, (k4_xboole_0(X1,k3_xboole_0(k2_xboole_0(X2,X1),X3))=k4_xboole_0(X1,X3)), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_100, plain, (r1_tarski(k3_xboole_0(X1,X2),k4_xboole_0(X2,k4_xboole_0(X3,X1)))), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_101, plain, (k3_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_65]), c_0_46])).
fof(c_0_102, negated_conjecture, k4_xboole_0(esk4_0,k4_xboole_0(esk4_0,esk5_0))!=k3_xboole_0(esk4_0,esk5_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_94])])])).
cnf(c_0_103, plain, (k3_xboole_0(X1,X2)=X2|~r1_tarski(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_97])])).
cnf(c_0_104, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k3_xboole_0(k2_xboole_0(X3,X1),X2))), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_105, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X3,X1)))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_100]), c_0_51]), c_0_101])).
cnf(c_0_106, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(X3,k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_46, c_0_51])).
cnf(c_0_107, plain, (k3_xboole_0(X1,k3_xboole_0(k2_xboole_0(X2,X1),X3))=k3_xboole_0(X1,X3)), inference(spm,[status(thm)],[c_0_51, c_0_91])).
cnf(c_0_108, negated_conjecture, (k4_xboole_0(esk4_0,k4_xboole_0(esk4_0,esk5_0))!=k3_xboole_0(esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_109, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_51]), c_0_105]), c_0_106]), c_0_107])).
cnf(c_0_110, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_108, c_0_109])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 111
# Proof object clause steps            : 71
# Proof object formula steps           : 40
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 39
# Proof object simplifying inferences  : 28
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 20
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 31
# Processed clauses                    : 13792
# ...of these trivial                  : 5332
# ...subsumed                          : 6531
# ...remaining for further processing  : 1929
# Other redundant clauses eliminated   : 1351
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 153
# Generated clauses                    : 719941
# ...of the previous two non-trivial   : 307428
# Contextual simplify-reflections      : 5
# Paramodulations                      : 718576
# Factorizations                       : 14
# Equation resolutions                 : 1351
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1769
#    Positive orientable unit clauses  : 1485
#    Positive unorientable unit clauses: 5
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 276
# Current number of unprocessed clauses: 292772
# ...number of literals in the above   : 424587
# Current number of archived formulas  : 0
# Current number of archived clauses   : 157
# Clause-clause subsumption calls (NU) : 45484
# Rec. Clause-clause subsumption calls : 39633
# Non-unit clause-clause subsumptions  : 3847
# Unit Clause-clause subsumption calls : 551
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12451
# BW rewrite match successes           : 331
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5969656

# -------------------------------------------------
# User time                : 3.996 s
# System time              : 0.086 s
# Total time               : 4.082 s
# Maximum resident set size: 251440 pages
