# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4_yellow19, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((v13_waybel_0(X3,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))))))=>(r2_waybel_7(X1,X3,X2)<=>r1_tarski(k1_yellow19(X1,X2),X3))))), file('data/MPTP/proofs/5e630a244e044b5bbe1d171a69924fb7.E_input', t4_yellow19)).
fof(dt_k1_yellow19, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k1_yellow19(X1,X2),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))), file('data/MPTP/proofs/5e630a244e044b5bbe1d171a69924fb7.E_input', dt_k1_yellow19)).
fof(fc1_yellow19, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,u1_struct_0(X1)))=>(~(v1_xboole_0(k1_yellow19(X1,X2)))&(v1_subset_1(k1_yellow19(X1,X2),u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(k1_yellow19(X1,X2),k3_yellow_1(k2_struct_0(X1)))&v13_waybel_0(k1_yellow19(X1,X2),k3_yellow_1(k2_struct_0(X1))))))), file('data/MPTP/proofs/5e630a244e044b5bbe1d171a69924fb7.E_input', fc1_yellow19)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/5e630a244e044b5bbe1d171a69924fb7.E_input', d10_xboole_0)).
fof(t5_yellow19, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r2_waybel_7(X1,k1_yellow19(X1,X2),X2))), file('data/MPTP/proofs/5e630a244e044b5bbe1d171a69924fb7.E_input', t5_yellow19)).
fof(c_0_5, plain, ![X4, X5, X6]:((~r2_waybel_7(X4,X6,X5)|r1_tarski(k1_yellow19(X4,X5),X6)|(~v13_waybel_0(X6,k3_yellow_1(k2_struct_0(X4)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X4))))))|~m1_subset_1(X5,u1_struct_0(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))&(~r1_tarski(k1_yellow19(X4,X5),X6)|r2_waybel_7(X4,X6,X5)|(~v13_waybel_0(X6,k3_yellow_1(k2_struct_0(X4)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X4))))))|~m1_subset_1(X5,u1_struct_0(X4))|(v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_yellow19])])])])])).
fof(c_0_6, plain, ![X7, X8]:(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))|~m1_subset_1(X8,u1_struct_0(X7))|m1_subset_1(k1_yellow19(X7,X8),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X7)))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_yellow19])])])).
fof(c_0_7, plain, ![X9, X10]:((~v1_xboole_0(k1_yellow19(X9,X10))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))|~m1_subset_1(X10,u1_struct_0(X9))))&((v1_subset_1(k1_yellow19(X9,X10),u1_struct_0(k3_yellow_1(k2_struct_0(X9))))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))|~m1_subset_1(X10,u1_struct_0(X9))))&((v2_waybel_0(k1_yellow19(X9,X10),k3_yellow_1(k2_struct_0(X9)))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))|~m1_subset_1(X10,u1_struct_0(X9))))&(v13_waybel_0(k1_yellow19(X9,X10),k3_yellow_1(k2_struct_0(X9)))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))|~m1_subset_1(X10,u1_struct_0(X9))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_yellow19])])])])).
fof(c_0_8, plain, ![X11, X12]:(((r1_tarski(X11,X12)|X11!=X12)&(r1_tarski(X12,X11)|X11!=X12))&(~r1_tarski(X11,X12)|~r1_tarski(X12,X11)|X11=X12)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_9, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r2_waybel_7(X1,k1_yellow19(X1,X2),X2)))), inference(assume_negation,[status(cth)],[t5_yellow19])).
cnf(c_0_10, plain, (r2_waybel_7(X1,X3,X2)|v2_struct_0(X1)|~r1_tarski(k1_yellow19(X1,X2),X3)|~v13_waybel_0(X3,k3_yellow_1(k2_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11, plain, (v2_struct_0(X1)|m1_subset_1(k1_yellow19(X1,X2),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12, plain, (v13_waybel_0(k1_yellow19(X1,X2),k3_yellow_1(k2_struct_0(X1)))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_14, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&~r2_waybel_7(esk1_0,k1_yellow19(esk1_0,esk2_0),esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])).
cnf(c_0_15, plain, (r2_waybel_7(X1,k1_yellow19(X1,X2),X3)|v2_struct_0(X1)|~r1_tarski(k1_yellow19(X1,X3),k1_yellow19(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])).
cnf(c_0_16, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_13])).
cnf(c_0_17, negated_conjecture, (~r2_waybel_7(esk1_0,k1_yellow19(esk1_0,esk2_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (r2_waybel_7(X1,k1_yellow19(X1,X2),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_19, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19]), c_0_20]), c_0_21])]), c_0_22]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 24
# Proof object clause steps            : 13
# Proof object formula steps           : 11
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 7
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 15
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 15
# Processed clauses                    : 32
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 32
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 6
# ...of the previous two non-trivial   : 4
# Contextual simplify-reflections      : 1
# Paramodulations                      : 4
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 16
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 10
# Current number of unprocessed clauses: 1
# ...number of literals in the above   : 8
# Current number of archived formulas  : 0
# Current number of archived clauses   : 14
# Clause-clause subsumption calls (NU) : 80
# Rec. Clause-clause subsumption calls : 1
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1548

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.003 s
# Total time               : 0.009 s
# Maximum resident set size: 28656 pages
