# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l3_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(X1,k1_tarski(X2))<=>(X1=k1_xboole_0|X1=k1_tarski(X2))), file('data/MPTP/proofs/5e67df629d734fb48efd249f356a0845.E_input', l3_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/5e67df629d734fb48efd249f356a0845.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/5e67df629d734fb48efd249f356a0845.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/5e67df629d734fb48efd249f356a0845.E_input', t71_enumset1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/5e67df629d734fb48efd249f356a0845.E_input', d1_tarski)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/5e67df629d734fb48efd249f356a0845.E_input', d3_tarski)).
fof(t11_setfam_1, conjecture, ![X1]:k1_setfam_1(k1_tarski(X1))=X1, file('data/MPTP/proofs/5e67df629d734fb48efd249f356a0845.E_input', t11_setfam_1)).
fof(d1_setfam_1, axiom, ![X1, X2]:((~(X1=k1_xboole_0)=>(X2=k1_setfam_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>![X4]:(r2_hidden(X4,X1)=>r2_hidden(X3,X4)))))&(X1=k1_xboole_0=>(X2=k1_setfam_1(X1)<=>X2=k1_xboole_0))), file('data/MPTP/proofs/5e67df629d734fb48efd249f356a0845.E_input', d1_setfam_1)).
fof(t64_zfmisc_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k4_xboole_0(X2,k1_tarski(X3)))<=>(r2_hidden(X1,X2)&~(X1=X3))), file('data/MPTP/proofs/5e67df629d734fb48efd249f356a0845.E_input', t64_zfmisc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/5e67df629d734fb48efd249f356a0845.E_input', d10_xboole_0)).
fof(c_0_10, plain, ![X5, X6]:((~r1_tarski(X5,k1_tarski(X6))|(X5=k1_xboole_0|X5=k1_tarski(X6)))&((X5!=k1_xboole_0|r1_tarski(X5,k1_tarski(X6)))&(X5!=k1_tarski(X6)|r1_tarski(X5,k1_tarski(X6))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_zfmisc_1])])])).
fof(c_0_11, plain, ![X7]:k2_tarski(X7,X7)=k1_tarski(X7), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_12, plain, ![X8, X9]:k1_enumset1(X8,X8,X9)=k2_tarski(X8,X9), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_13, plain, ![X10, X11, X12]:k2_enumset1(X10,X10,X11,X12)=k1_enumset1(X10,X11,X12), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_14, plain, ![X13, X14, X15, X16, X17, X18]:(((~r2_hidden(X15,X14)|X15=X13|X14!=k1_tarski(X13))&(X16!=X13|r2_hidden(X16,X14)|X14!=k1_tarski(X13)))&((~r2_hidden(esk1_2(X17,X18),X18)|esk1_2(X17,X18)!=X17|X18=k1_tarski(X17))&(r2_hidden(esk1_2(X17,X18),X18)|esk1_2(X17,X18)=X17|X18=k1_tarski(X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_15, plain, ![X20, X21, X22, X23, X24]:((~r1_tarski(X20,X21)|(~r2_hidden(X22,X20)|r2_hidden(X22,X21)))&((r2_hidden(esk2_2(X23,X24),X23)|r1_tarski(X23,X24))&(~r2_hidden(esk2_2(X23,X24),X24)|r1_tarski(X23,X24)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_16, plain, (r1_tarski(X1,k1_tarski(X2))|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (r1_tarski(X1,k2_enumset1(X2,X2,X2,X2))|X1!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_17]), c_0_18]), c_0_19])).
fof(c_0_24, negated_conjecture, ~(![X1]:k1_setfam_1(k1_tarski(X1))=X1), inference(assume_negation,[status(cth)],[t11_setfam_1])).
cnf(c_0_25, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_enumset1(X2,X2,X2,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_17]), c_0_18]), c_0_19])).
cnf(c_0_26, plain, (X1=X3|X2!=k2_enumset1(X3,X3,X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_17]), c_0_18]), c_0_19])).
cnf(c_0_27, plain, (r2_hidden(X1,k2_enumset1(X2,X2,X2,X2))|X3!=k1_xboole_0|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_28, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_29, negated_conjecture, k1_setfam_1(k1_tarski(esk3_0))!=esk3_0, inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).
fof(c_0_30, plain, ![X27, X28, X29, X30, X31, X33, X36, X37, X38, X39]:((((~r2_hidden(X29,X28)|(~r2_hidden(X30,X27)|r2_hidden(X29,X30))|X28!=k1_setfam_1(X27)|X27=k1_xboole_0)&((r2_hidden(esk4_3(X27,X28,X31),X27)|r2_hidden(X31,X28)|X28!=k1_setfam_1(X27)|X27=k1_xboole_0)&(~r2_hidden(X31,esk4_3(X27,X28,X31))|r2_hidden(X31,X28)|X28!=k1_setfam_1(X27)|X27=k1_xboole_0)))&(((r2_hidden(esk6_2(X27,X33),X27)|~r2_hidden(esk5_2(X27,X33),X33)|X33=k1_setfam_1(X27)|X27=k1_xboole_0)&(~r2_hidden(esk5_2(X27,X33),esk6_2(X27,X33))|~r2_hidden(esk5_2(X27,X33),X33)|X33=k1_setfam_1(X27)|X27=k1_xboole_0))&(r2_hidden(esk5_2(X27,X33),X33)|(~r2_hidden(X36,X27)|r2_hidden(esk5_2(X27,X33),X36))|X33=k1_setfam_1(X27)|X27=k1_xboole_0)))&((X38!=k1_setfam_1(X37)|X38=k1_xboole_0|X37!=k1_xboole_0)&(X39!=k1_xboole_0|X39=k1_setfam_1(X37)|X37!=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_setfam_1])])])])])])])).
cnf(c_0_31, plain, (r2_hidden(X1,X2)|X2!=k2_enumset1(X1,X1,X1,X1)), inference(er,[status(thm)],[c_0_25])).
cnf(c_0_32, plain, (X1=X2|~r2_hidden(X1,k2_enumset1(X2,X2,X2,X2))), inference(er,[status(thm)],[c_0_26])).
cnf(c_0_33, plain, (r2_hidden(esk2_2(X1,X2),k2_enumset1(X3,X3,X3,X3))|r1_tarski(X1,X2)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34, negated_conjecture, (k1_setfam_1(k1_tarski(esk3_0))!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_35, plain, (r2_hidden(esk5_2(X1,X2),X2)|r2_hidden(esk5_2(X1,X2),X3)|X2=k1_setfam_1(X1)|X1=k1_xboole_0|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_36, plain, (r2_hidden(X1,k2_enumset1(X1,X1,X1,X1))), inference(er,[status(thm)],[c_0_31])).
cnf(c_0_37, plain, (esk2_2(X1,X2)=X3|r1_tarski(X1,X2)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_32, c_0_33])).
fof(c_0_38, plain, ![X40, X41, X42]:(((r2_hidden(X40,X41)|~r2_hidden(X40,k4_xboole_0(X41,k1_tarski(X42))))&(X40!=X42|~r2_hidden(X40,k4_xboole_0(X41,k1_tarski(X42)))))&(~r2_hidden(X40,X41)|X40=X42|r2_hidden(X40,k4_xboole_0(X41,k1_tarski(X42))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t64_zfmisc_1])])])])).
cnf(c_0_39, negated_conjecture, (k1_setfam_1(k2_enumset1(esk3_0,esk3_0,esk3_0,esk3_0))!=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_17]), c_0_18]), c_0_19])).
cnf(c_0_40, plain, (X1=k1_setfam_1(k2_enumset1(X2,X2,X2,X2))|k2_enumset1(X2,X2,X2,X2)=k1_xboole_0|r2_hidden(esk5_2(k2_enumset1(X2,X2,X2,X2),X1),X1)|r2_hidden(esk5_2(k2_enumset1(X2,X2,X2,X2),X1),X2)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_41, plain, (X1=X2|r1_tarski(X3,X4)|X3!=k1_xboole_0), inference(spm,[status(thm)],[c_0_37, c_0_37])).
cnf(c_0_42, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,k1_tarski(X3)))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_43, negated_conjecture, (k2_enumset1(esk3_0,esk3_0,esk3_0,esk3_0)=k1_xboole_0|r2_hidden(esk5_2(k2_enumset1(esk3_0,esk3_0,esk3_0,esk3_0),X1),esk3_0)|r2_hidden(esk5_2(k2_enumset1(esk3_0,esk3_0,esk3_0,esk3_0),X1),X1)|X1!=esk3_0), inference(spm,[status(thm)],[c_0_39, c_0_40])).
fof(c_0_44, plain, ![X43, X44]:(((r1_tarski(X43,X44)|X43!=X44)&(r1_tarski(X44,X43)|X43!=X44))&(~r1_tarski(X43,X44)|~r1_tarski(X44,X43)|X43=X44)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_45, plain, (X1=X2|r1_tarski(k1_xboole_0,X3)), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_46, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,k2_enumset1(X3,X3,X3,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_17]), c_0_18]), c_0_19])).
cnf(c_0_47, plain, (r2_hidden(esk6_2(X1,X2),X1)|X2=k1_setfam_1(X1)|X1=k1_xboole_0|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_48, negated_conjecture, (k2_enumset1(esk3_0,esk3_0,esk3_0,esk3_0)=k1_xboole_0|r2_hidden(esk5_2(k2_enumset1(esk3_0,esk3_0,esk3_0,esk3_0),esk3_0),esk3_0)), inference(er,[status(thm)],[c_0_43])).
cnf(c_0_49, plain, (X1!=X2|~r2_hidden(X1,k4_xboole_0(X3,k1_tarski(X2)))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_51, negated_conjecture, (r1_tarski(k1_xboole_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_45]), c_0_45])).
cnf(c_0_52, plain, (r1_tarski(X1,X2)|~r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_53, plain, (r2_hidden(esk2_2(k4_xboole_0(X1,k2_enumset1(X2,X2,X2,X2)),X3),X1)|r1_tarski(k4_xboole_0(X1,k2_enumset1(X2,X2,X2,X2)),X3)), inference(spm,[status(thm)],[c_0_46, c_0_28])).
cnf(c_0_54, negated_conjecture, (k2_enumset1(esk3_0,esk3_0,esk3_0,esk3_0)=k1_xboole_0|r2_hidden(esk6_2(k2_enumset1(esk3_0,esk3_0,esk3_0,esk3_0),esk3_0),k2_enumset1(esk3_0,esk3_0,esk3_0,esk3_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_39])).
cnf(c_0_55, plain, (X1!=X2|~r2_hidden(X1,k4_xboole_0(X3,k2_enumset1(X2,X2,X2,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_17]), c_0_18]), c_0_19])).
cnf(c_0_56, negated_conjecture, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_57, plain, (r1_tarski(k4_xboole_0(X1,k2_enumset1(X2,X2,X2,X2)),X1)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_58, plain, (X2=k1_setfam_1(X1)|X1=k1_xboole_0|~r2_hidden(esk5_2(X1,X2),esk6_2(X1,X2))|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_59, negated_conjecture, (esk6_2(k2_enumset1(esk3_0,esk3_0,esk3_0,esk3_0),esk3_0)=esk3_0|k2_enumset1(esk3_0,esk3_0,esk3_0,esk3_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_32, c_0_54])).
cnf(c_0_60, plain, (~r2_hidden(X1,k4_xboole_0(X2,k2_enumset1(X1,X1,X1,X1)))), inference(er,[status(thm)],[c_0_55])).
cnf(c_0_61, negated_conjecture, (k4_xboole_0(k1_xboole_0,k2_enumset1(X1,X1,X1,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_62, negated_conjecture, (k2_enumset1(esk3_0,esk3_0,esk3_0,esk3_0)=k1_xboole_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_39]), c_0_48])).
cnf(c_0_63, negated_conjecture, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_64, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_62]), c_0_63]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 65
# Proof object clause steps            : 44
# Proof object formula steps           : 21
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 25
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 27
# Processed clauses                    : 453
# ...of these trivial                  : 21
# ...subsumed                          : 183
# ...remaining for further processing  : 249
# Other redundant clauses eliminated   : 148
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 9
# Backward-rewritten                   : 11
# Generated clauses                    : 6553
# ...of the previous two non-trivial   : 6010
# Contextual simplify-reflections      : 9
# Paramodulations                      : 6344
# Factorizations                       : 25
# Equation resolutions                 : 184
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 199
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 190
# Current number of unprocessed clauses: 5361
# ...number of literals in the above   : 36563
# Current number of archived formulas  : 0
# Current number of archived clauses   : 49
# Clause-clause subsumption calls (NU) : 4398
# Rec. Clause-clause subsumption calls : 1571
# Non-unit clause-clause subsumptions  : 152
# Unit Clause-clause subsumption calls : 46
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 118261

# -------------------------------------------------
# User time                : 0.110 s
# System time              : 0.003 s
# Total time               : 0.113 s
# Maximum resident set size: 28568 pages
