# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.007 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t44_zfmisc_1, conjecture, ![X1, X2, X3]:~((k1_tarski(X1)=k2_xboole_0(X2,X3)&(~(X2=X3)&(~(X2=k1_xboole_0)&~(X3=k1_xboole_0))))), file('data/MPTP/proofs/5e9a92dd08724271a644fc916f9d8daf.E_input', t44_zfmisc_1)).
fof(t43_zfmisc_1, axiom, ![X1, X2, X3]:~((k1_tarski(X1)=k2_xboole_0(X2,X3)&(~((X2=k1_tarski(X1)&X3=k1_tarski(X1)))&(~((X2=k1_xboole_0&X3=k1_tarski(X1)))&~((X2=k1_tarski(X1)&X3=k1_xboole_0)))))), file('data/MPTP/proofs/5e9a92dd08724271a644fc916f9d8daf.E_input', t43_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/5e9a92dd08724271a644fc916f9d8daf.E_input', t69_enumset1)).
fof(c_0_3, negated_conjecture, ~(![X1, X2, X3]:~((k1_tarski(X1)=k2_xboole_0(X2,X3)&(~(X2=X3)&(~(X2=k1_xboole_0)&~(X3=k1_xboole_0)))))), inference(assume_negation,[status(cth)],[t44_zfmisc_1])).
fof(c_0_4, plain, ![X7, X8, X9]:((((X8=k1_tarski(X7)|X8=k1_xboole_0|X8=k1_tarski(X7)|k1_tarski(X7)!=k2_xboole_0(X8,X9))&(X9=k1_xboole_0|X8=k1_xboole_0|X8=k1_tarski(X7)|k1_tarski(X7)!=k2_xboole_0(X8,X9)))&((X8=k1_tarski(X7)|X9=k1_tarski(X7)|X8=k1_tarski(X7)|k1_tarski(X7)!=k2_xboole_0(X8,X9))&(X9=k1_xboole_0|X9=k1_tarski(X7)|X8=k1_tarski(X7)|k1_tarski(X7)!=k2_xboole_0(X8,X9))))&(((X8=k1_tarski(X7)|X8=k1_xboole_0|X9=k1_tarski(X7)|k1_tarski(X7)!=k2_xboole_0(X8,X9))&(X9=k1_xboole_0|X8=k1_xboole_0|X9=k1_tarski(X7)|k1_tarski(X7)!=k2_xboole_0(X8,X9)))&((X8=k1_tarski(X7)|X9=k1_tarski(X7)|X9=k1_tarski(X7)|k1_tarski(X7)!=k2_xboole_0(X8,X9))&(X9=k1_xboole_0|X9=k1_tarski(X7)|X9=k1_tarski(X7)|k1_tarski(X7)!=k2_xboole_0(X8,X9))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_zfmisc_1])])])).
fof(c_0_5, plain, ![X10]:k2_tarski(X10,X10)=k1_tarski(X10), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_6, negated_conjecture, (k1_tarski(esk1_0)=k2_xboole_0(esk2_0,esk3_0)&(esk2_0!=esk3_0&(esk2_0!=k1_xboole_0&esk3_0!=k1_xboole_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_3])])])])).
cnf(c_0_7, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|X1=k1_tarski(X2)|k1_tarski(X2)!=k2_xboole_0(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, negated_conjecture, (k1_tarski(esk1_0)=k2_xboole_0(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (X1=k1_xboole_0|X1=k2_tarski(X2,X2)|k2_tarski(X2,X2)!=k2_xboole_0(X3,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[c_0_7]), c_0_8]), c_0_8])).
cnf(c_0_11, negated_conjecture, (k2_tarski(esk1_0,esk1_0)=k2_xboole_0(esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_9, c_0_8])).
cnf(c_0_12, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13, negated_conjecture, (k2_tarski(X1,X1)=esk3_0|k2_tarski(X1,X1)!=k2_tarski(esk1_0,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])).
cnf(c_0_14, plain, (X1=k1_tarski(X2)|X1=k1_xboole_0|X1=k1_tarski(X2)|k1_tarski(X2)!=k2_xboole_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_15, negated_conjecture, (k2_tarski(esk1_0,esk1_0)=esk3_0), inference(er,[status(thm)],[c_0_13])).
cnf(c_0_16, plain, (X1=k1_xboole_0|X1=k2_tarski(X2,X2)|k2_tarski(X2,X2)!=k2_xboole_0(X1,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[c_0_14]), c_0_8]), c_0_8])).
cnf(c_0_17, negated_conjecture, (k2_xboole_0(esk2_0,esk3_0)=esk3_0), inference(rw,[status(thm)],[c_0_11, c_0_15])).
cnf(c_0_18, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_19, negated_conjecture, (k2_tarski(X1,X1)=esk2_0|k2_tarski(X1,X1)!=esk3_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])).
cnf(c_0_20, negated_conjecture, (esk2_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_21, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_15]), c_0_20]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 22
# Proof object clause steps            : 15
# Proof object formula steps           : 7
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 7
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 4
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 13
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 12
# Processed clauses                    : 22
# ...of these trivial                  : 0
# ...subsumed                          : 5
# ...remaining for further processing  : 17
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 5
# ...of the previous two non-trivial   : 4
# Contextual simplify-reflections      : 0
# Paramodulations                      : 4
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 8
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 3
# Current number of unprocessed clauses: 1
# ...number of literals in the above   : 3
# Current number of archived formulas  : 0
# Current number of archived clauses   : 10
# Clause-clause subsumption calls (NU) : 10
# Rec. Clause-clause subsumption calls : 10
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 472

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.000 s
# Total time               : 0.008 s
# Maximum resident set size: 28620 pages
