# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t42_tex_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v2_tex_2(X2,X1)=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,X2)=>k9_subset_1(u1_struct_0(X1),X2,k2_pre_topc(X1,k6_domain_1(u1_struct_0(X1),X3)))=k6_domain_1(u1_struct_0(X1),X3)))))), file('data/MPTP/proofs/5ef1a0ba0e9841a284f0e3fb47529c41.E_input', t42_tex_2)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/5ef1a0ba0e9841a284f0e3fb47529c41.E_input', t3_subset)).
fof(l1_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/5ef1a0ba0e9841a284f0e3fb47529c41.E_input', l1_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/5ef1a0ba0e9841a284f0e3fb47529c41.E_input', t69_enumset1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/5ef1a0ba0e9841a284f0e3fb47529c41.E_input', d3_tarski)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/5ef1a0ba0e9841a284f0e3fb47529c41.E_input', t1_xboole_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/5ef1a0ba0e9841a284f0e3fb47529c41.E_input', d1_xboole_0)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/5ef1a0ba0e9841a284f0e3fb47529c41.E_input', redefinition_k6_domain_1)).
fof(t41_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v2_tex_2(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r1_tarski(X3,X2)=>k9_subset_1(u1_struct_0(X1),X2,k2_pre_topc(X1,X3))=X3))))), file('data/MPTP/proofs/5ef1a0ba0e9841a284f0e3fb47529c41.E_input', t41_tex_2)).
fof(c_0_9, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v2_tex_2(X2,X1)=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,X2)=>k9_subset_1(u1_struct_0(X1),X2,k2_pre_topc(X1,k6_domain_1(u1_struct_0(X1),X3)))=k6_domain_1(u1_struct_0(X1),X3))))))), inference(assume_negation,[status(cth)],[t42_tex_2])).
fof(c_0_10, plain, ![X7, X8]:((~m1_subset_1(X7,k1_zfmisc_1(X8))|r1_tarski(X7,X8))&(~r1_tarski(X7,X8)|m1_subset_1(X7,k1_zfmisc_1(X8)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_11, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(v2_tex_2(esk2_0,esk1_0)&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(r2_hidden(esk3_0,esk2_0)&k9_subset_1(u1_struct_0(esk1_0),esk2_0,k2_pre_topc(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))!=k6_domain_1(u1_struct_0(esk1_0),esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])).
fof(c_0_12, plain, ![X9, X10]:((~r1_tarski(k1_tarski(X9),X10)|r2_hidden(X9,X10))&(~r2_hidden(X9,X10)|r1_tarski(k1_tarski(X9),X10))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])).
fof(c_0_13, plain, ![X11]:k2_tarski(X11,X11)=k1_tarski(X11), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_14, plain, ![X12, X13, X14, X15, X16]:((~r1_tarski(X12,X13)|(~r2_hidden(X14,X12)|r2_hidden(X14,X13)))&((r2_hidden(esk4_2(X15,X16),X15)|r1_tarski(X15,X16))&(~r2_hidden(esk4_2(X15,X16),X16)|r1_tarski(X15,X16)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_15, plain, ![X18, X19, X20]:(~r1_tarski(X18,X19)|~r1_tarski(X19,X20)|r1_tarski(X18,X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
cnf(c_0_16, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, plain, (r2_hidden(X1,X2)|~r1_tarski(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, plain, (r1_tarski(X1,X2)|~r2_hidden(esk4_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, negated_conjecture, (r1_tarski(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_24, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_25, plain, ![X21, X22, X23]:((~v1_xboole_0(X21)|~r2_hidden(X22,X21))&(r2_hidden(esk5_1(X23),X23)|v1_xboole_0(X23))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_26, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_tarski(X1,X1),X2)), inference(rw,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_27, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_28, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_29, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~r1_tarski(X1,esk2_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_30, plain, (r1_tarski(k2_tarski(X1,X1),X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_24, c_0_19])).
cnf(c_0_31, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, plain, (r2_hidden(X1,k2_tarski(X1,X1))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_33, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r1_tarski(X2,esk2_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k2_tarski(X3,X3))|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_28, c_0_30])).
cnf(c_0_35, plain, (r2_hidden(esk5_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_36, plain, (~v1_xboole_0(k2_tarski(X1,X1))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
fof(c_0_37, plain, ![X25, X26]:(v1_xboole_0(X25)|~m1_subset_1(X26,X25)|k6_domain_1(X25,X26)=k1_tarski(X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_38, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k2_tarski(X2,X2))|~r2_hidden(X2,esk2_0)), inference(spm,[status(thm)],[c_0_33, c_0_30])).
cnf(c_0_39, plain, (r2_hidden(esk5_1(k2_tarski(X1,X1)),X2)|~r2_hidden(X1,X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])).
cnf(c_0_40, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_41, plain, ![X27, X28, X29]:((~v2_tex_2(X28,X27)|(~m1_subset_1(X29,k1_zfmisc_1(u1_struct_0(X27)))|(~r1_tarski(X29,X28)|k9_subset_1(u1_struct_0(X27),X28,k2_pre_topc(X27,X29))=X29))|~m1_subset_1(X28,k1_zfmisc_1(u1_struct_0(X27)))|(v2_struct_0(X27)|(~v2_pre_topc(X27)|~l1_pre_topc(X27))))&((m1_subset_1(esk6_2(X27,X28),k1_zfmisc_1(u1_struct_0(X27)))|v2_tex_2(X28,X27)|~m1_subset_1(X28,k1_zfmisc_1(u1_struct_0(X27)))|(v2_struct_0(X27)|(~v2_pre_topc(X27)|~l1_pre_topc(X27))))&((r1_tarski(esk6_2(X27,X28),X28)|v2_tex_2(X28,X27)|~m1_subset_1(X28,k1_zfmisc_1(u1_struct_0(X27)))|(v2_struct_0(X27)|(~v2_pre_topc(X27)|~l1_pre_topc(X27))))&(k9_subset_1(u1_struct_0(X27),X28,k2_pre_topc(X27,esk6_2(X27,X28)))!=esk6_2(X27,X28)|v2_tex_2(X28,X27)|~m1_subset_1(X28,k1_zfmisc_1(u1_struct_0(X27)))|(v2_struct_0(X27)|(~v2_pre_topc(X27)|~l1_pre_topc(X27))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t41_tex_2])])])])])])).
cnf(c_0_42, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_28, c_0_23])).
cnf(c_0_43, plain, (r1_tarski(k2_tarski(X1,X1),X2)|r2_hidden(esk4_2(k2_tarski(X1,X1),X2),X3)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_34, c_0_21])).
cnf(c_0_44, negated_conjecture, (r2_hidden(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_45, negated_conjecture, (r2_hidden(esk5_1(k2_tarski(X1,X1)),u1_struct_0(esk1_0))|~r2_hidden(X1,k2_tarski(X2,X2))|~r2_hidden(X2,esk2_0)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_46, plain, (k2_tarski(X2,X2)=k6_domain_1(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_40, c_0_19])).
cnf(c_0_47, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_48, plain, (k9_subset_1(u1_struct_0(X2),X1,k2_pre_topc(X2,X3))=X3|v2_struct_0(X2)|~v2_tex_2(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~r1_tarski(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_49, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_50, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~r2_hidden(esk4_2(X1,u1_struct_0(esk1_0)),esk2_0)), inference(spm,[status(thm)],[c_0_20, c_0_42])).
cnf(c_0_51, negated_conjecture, (r1_tarski(k2_tarski(esk3_0,esk3_0),X1)|r2_hidden(esk4_2(k2_tarski(esk3_0,esk3_0),X1),esk2_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk5_1(k2_tarski(X1,X1)),u1_struct_0(esk1_0))|~r2_hidden(X1,k2_tarski(esk3_0,esk3_0))), inference(spm,[status(thm)],[c_0_45, c_0_44])).
cnf(c_0_53, negated_conjecture, (k9_subset_1(u1_struct_0(esk1_0),esk2_0,k2_pre_topc(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))!=k6_domain_1(u1_struct_0(esk1_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_54, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk3_0)=k2_tarski(esk3_0,esk3_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_55, plain, (k9_subset_1(u1_struct_0(X1),X2,k2_pre_topc(X1,X3))=X3|v2_struct_0(X1)|~r1_tarski(X3,u1_struct_0(X1))|~r1_tarski(X3,X2)|~v2_tex_2(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_56, negated_conjecture, (r1_tarski(k2_tarski(esk3_0,esk3_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_57, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_58, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_59, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_60, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))|~r2_hidden(X1,k2_tarski(esk3_0,esk3_0))), inference(spm,[status(thm)],[c_0_31, c_0_52])).
cnf(c_0_61, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|k9_subset_1(u1_struct_0(esk1_0),esk2_0,k2_pre_topc(esk1_0,k2_tarski(esk3_0,esk3_0)))!=k2_tarski(esk3_0,esk3_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_62, negated_conjecture, (k9_subset_1(u1_struct_0(esk1_0),X1,k2_pre_topc(esk1_0,k2_tarski(esk3_0,esk3_0)))=k2_tarski(esk3_0,esk3_0)|~r1_tarski(k2_tarski(esk3_0,esk3_0),X1)|~v2_tex_2(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_63, negated_conjecture, (r1_tarski(k2_tarski(esk3_0,esk3_0),esk2_0)), inference(spm,[status(thm)],[c_0_20, c_0_51])).
cnf(c_0_64, negated_conjecture, (v2_tex_2(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_65, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_35]), c_0_36])).
cnf(c_0_66, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63]), c_0_64]), c_0_17])]), c_0_65]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 67
# Proof object clause steps            : 48
# Proof object formula steps           : 19
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 24
# Proof object simplifying inferences  : 14
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 23
# Processed clauses                    : 533
# ...of these trivial                  : 1
# ...subsumed                          : 260
# ...remaining for further processing  : 272
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 1
# Generated clauses                    : 1266
# ...of the previous two non-trivial   : 1191
# Contextual simplify-reflections      : 1
# Paramodulations                      : 1265
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 244
#    Positive orientable unit clauses  : 23
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 216
# Current number of unprocessed clauses: 699
# ...number of literals in the above   : 2700
# Current number of archived formulas  : 0
# Current number of archived clauses   : 29
# Clause-clause subsumption calls (NU) : 10447
# Rec. Clause-clause subsumption calls : 5557
# Non-unit clause-clause subsumptions  : 208
# Unit Clause-clause subsumption calls : 148
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 23
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 33678

# -------------------------------------------------
# User time                : 0.028 s
# System time              : 0.003 s
# Total time               : 0.032 s
# Maximum resident set size: 28788 pages
