# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t24_yellow_0, conjecture, ![X1]:((v3_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X2=k13_lattice3(X1,X2,X3)<=>r1_orders_2(X1,X3,X2))))), file('data/MPTP/proofs/5ef93ce086074fe48fcdfe4b53c2be71.E_input', t24_yellow_0)).
fof(t22_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k13_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))))))), file('data/MPTP/proofs/5ef93ce086074fe48fcdfe4b53c2be71.E_input', t22_yellow_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/5ef93ce086074fe48fcdfe4b53c2be71.E_input', t24_orders_2)).
fof(cc1_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v1_lattice3(X1)=>~(v2_struct_0(X1)))), file('data/MPTP/proofs/5ef93ce086074fe48fcdfe4b53c2be71.E_input', cc1_lattice3)).
fof(c_0_4, negated_conjecture, ~(![X1]:((v3_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X2=k13_lattice3(X1,X2,X3)<=>r1_orders_2(X1,X3,X2)))))), inference(assume_negation,[status(cth)],[t24_yellow_0])).
fof(c_0_5, plain, ![X9, X10, X11, X12, X13]:(((r1_orders_2(X9,X10,X12)|X12!=k13_lattice3(X9,X10,X11)|~m1_subset_1(X12,u1_struct_0(X9))|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(~v5_orders_2(X9)|(~v1_lattice3(X9)|~l1_orders_2(X9))))&((r1_orders_2(X9,X11,X12)|X12!=k13_lattice3(X9,X10,X11)|~m1_subset_1(X12,u1_struct_0(X9))|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(~v5_orders_2(X9)|(~v1_lattice3(X9)|~l1_orders_2(X9))))&(~m1_subset_1(X13,u1_struct_0(X9))|(~r1_orders_2(X9,X10,X13)|~r1_orders_2(X9,X11,X13)|r1_orders_2(X9,X12,X13))|X12!=k13_lattice3(X9,X10,X11)|~m1_subset_1(X12,u1_struct_0(X9))|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(~v5_orders_2(X9)|(~v1_lattice3(X9)|~l1_orders_2(X9))))))&((m1_subset_1(esk4_4(X9,X10,X11,X12),u1_struct_0(X9))|~r1_orders_2(X9,X11,X12)|~r1_orders_2(X9,X10,X12)|X12=k13_lattice3(X9,X10,X11)|~m1_subset_1(X12,u1_struct_0(X9))|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(~v5_orders_2(X9)|(~v1_lattice3(X9)|~l1_orders_2(X9))))&(((r1_orders_2(X9,X10,esk4_4(X9,X10,X11,X12))|~r1_orders_2(X9,X11,X12)|~r1_orders_2(X9,X10,X12)|X12=k13_lattice3(X9,X10,X11)|~m1_subset_1(X12,u1_struct_0(X9))|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(~v5_orders_2(X9)|(~v1_lattice3(X9)|~l1_orders_2(X9))))&(r1_orders_2(X9,X11,esk4_4(X9,X10,X11,X12))|~r1_orders_2(X9,X11,X12)|~r1_orders_2(X9,X10,X12)|X12=k13_lattice3(X9,X10,X11)|~m1_subset_1(X12,u1_struct_0(X9))|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(~v5_orders_2(X9)|(~v1_lattice3(X9)|~l1_orders_2(X9)))))&(~r1_orders_2(X9,X12,esk4_4(X9,X10,X11,X12))|~r1_orders_2(X9,X11,X12)|~r1_orders_2(X9,X10,X12)|X12=k13_lattice3(X9,X10,X11)|~m1_subset_1(X12,u1_struct_0(X9))|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|(~v5_orders_2(X9)|(~v1_lattice3(X9)|~l1_orders_2(X9))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t22_yellow_0])])])])])).
fof(c_0_6, negated_conjecture, ((v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_lattice3(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&((esk2_0!=k13_lattice3(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk3_0,esk2_0))&(esk2_0=k13_lattice3(esk1_0,esk2_0,esk3_0)|r1_orders_2(esk1_0,esk3_0,esk2_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
cnf(c_0_7, plain, (r1_orders_2(X1,X2,X3)|X3!=k13_lattice3(X1,X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8, negated_conjecture, (esk2_0=k13_lattice3(esk1_0,esk2_0,esk3_0)|r1_orders_2(esk1_0,esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12, negated_conjecture, (v1_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_14, plain, ![X15, X16]:(v2_struct_0(X15)|(~v3_orders_2(X15)|~l1_orders_2(X15))|(~m1_subset_1(X16,u1_struct_0(X15))|r1_orders_2(X15,X16,X16))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
cnf(c_0_15, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk2_0)|r1_orders_2(esk1_0,esk3_0,X1)|X1!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_9]), c_0_10]), c_0_11]), c_0_12]), c_0_13])])).
cnf(c_0_16, plain, (X2=k13_lattice3(X1,X3,X4)|~r1_orders_2(X1,X2,esk4_4(X1,X3,X4,X2))|~r1_orders_2(X1,X4,X2)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_17, plain, (r1_orders_2(X1,X2,esk4_4(X1,X2,X3,X4))|X4=k13_lattice3(X1,X2,X3)|~r1_orders_2(X1,X3,X4)|~r1_orders_2(X1,X2,X4)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_18, plain, ![X17]:(~l1_orders_2(X17)|(~v1_lattice3(X17)|~v2_struct_0(X17))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattice3])])])).
cnf(c_0_19, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_21, negated_conjecture, (esk2_0!=k13_lattice3(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_22, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_15, c_0_9])).
cnf(c_0_23, plain, (k13_lattice3(X1,X2,X3)=X2|~r1_orders_2(X1,X3,X2)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_24, plain, (~l1_orders_2(X1)|~v1_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_9]), c_0_11]), c_0_20])])).
cnf(c_0_26, negated_conjecture, (k13_lattice3(esk1_0,esk2_0,esk3_0)!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_22])])).
cnf(c_0_27, negated_conjecture, (k13_lattice3(esk1_0,X1,X2)=X1|~r1_orders_2(esk1_0,X2,X1)|~r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_13]), c_0_11]), c_0_12])])).
cnf(c_0_28, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_11]), c_0_12])])).
cnf(c_0_29, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_22]), c_0_28]), c_0_10]), c_0_9])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 30
# Proof object clause steps            : 21
# Proof object formula steps           : 9
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 46
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 46
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 5
# Generated clauses                    : 34
# ...of the previous two non-trivial   : 19
# Contextual simplify-reflections      : 0
# Paramodulations                      : 25
# Factorizations                       : 0
# Equation resolutions                 : 9
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 24
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 14
# Current number of unprocessed clauses: 5
# ...number of literals in the above   : 45
# Current number of archived formulas  : 0
# Current number of archived clauses   : 22
# Clause-clause subsumption calls (NU) : 330
# Rec. Clause-clause subsumption calls : 40
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2640

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.003 s
# Total time               : 0.010 s
# Maximum resident set size: 28584 pages
