# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/5f20e94a25c24cd98ac4aa95def75f96.E_input', t7_xboole_1)).
fof(d6_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k3_relat_1(X1)=k2_xboole_0(k1_relat_1(X1),k2_relat_1(X1))), file('data/MPTP/proofs/5f20e94a25c24cd98ac4aa95def75f96.E_input', d6_relat_1)).
fof(d1_wellord2, axiom, ![X1, X2]:(v1_relat_1(X2)=>(X2=k1_wellord2(X1)<=>(k3_relat_1(X2)=X1&![X3, X4]:((r2_hidden(X3,X1)&r2_hidden(X4,X1))=>(r2_hidden(k4_tarski(X3,X4),X2)<=>r1_tarski(X3,X4)))))), file('data/MPTP/proofs/5f20e94a25c24cd98ac4aa95def75f96.E_input', d1_wellord2)).
fof(dt_k1_wellord2, axiom, ![X1]:v1_relat_1(k1_wellord2(X1)), file('data/MPTP/proofs/5f20e94a25c24cd98ac4aa95def75f96.E_input', dt_k1_wellord2)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/5f20e94a25c24cd98ac4aa95def75f96.E_input', d10_xboole_0)).
fof(t25_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>(r1_tarski(k1_relat_1(X1),k1_relat_1(X2))&r1_tarski(k2_relat_1(X1),k2_relat_1(X2)))))), file('data/MPTP/proofs/5f20e94a25c24cd98ac4aa95def75f96.E_input', t25_relat_1)).
fof(t71_relat_1, axiom, ![X1]:(k1_relat_1(k6_relat_1(X1))=X1&k2_relat_1(k6_relat_1(X1))=X1), file('data/MPTP/proofs/5f20e94a25c24cd98ac4aa95def75f96.E_input', t71_relat_1)).
fof(fc3_funct_1, axiom, ![X1]:(v1_relat_1(k6_relat_1(X1))&v1_funct_1(k6_relat_1(X1))), file('data/MPTP/proofs/5f20e94a25c24cd98ac4aa95def75f96.E_input', fc3_funct_1)).
fof(t73_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(![X3]:(r2_hidden(X3,X1)=>r2_hidden(k4_tarski(X3,X3),X2))=>r1_tarski(k6_relat_1(X1),X2))), file('data/MPTP/proofs/5f20e94a25c24cd98ac4aa95def75f96.E_input', t73_relat_1)).
fof(t88_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k7_relat_1(X2,X1),X2)), file('data/MPTP/proofs/5f20e94a25c24cd98ac4aa95def75f96.E_input', t88_relat_1)).
fof(t17_wellord1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k2_wellord1(X2,X1)=k7_relat_1(k8_relat_1(X1,X2),X1)), file('data/MPTP/proofs/5f20e94a25c24cd98ac4aa95def75f96.E_input', t17_wellord1)).
fof(dt_k8_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>v1_relat_1(k8_relat_1(X1,X2))), file('data/MPTP/proofs/5f20e94a25c24cd98ac4aa95def75f96.E_input', dt_k8_relat_1)).
fof(t30_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>k2_wellord1(X1,k3_relat_1(X1))=X1), file('data/MPTP/proofs/5f20e94a25c24cd98ac4aa95def75f96.E_input', t30_wellord1)).
fof(t117_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k8_relat_1(X1,X2),X2)), file('data/MPTP/proofs/5f20e94a25c24cd98ac4aa95def75f96.E_input', t117_relat_1)).
fof(t116_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k2_relat_1(k8_relat_1(X1,X2)),X1)), file('data/MPTP/proofs/5f20e94a25c24cd98ac4aa95def75f96.E_input', t116_relat_1)).
fof(t21_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>r1_tarski(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))), file('data/MPTP/proofs/5f20e94a25c24cd98ac4aa95def75f96.E_input', t21_relat_1)).
fof(t33_wellord2, conjecture, ![X1]:r1_tarski(k1_wellord2(X1),k2_zfmisc_1(X1,X1)), file('data/MPTP/proofs/5f20e94a25c24cd98ac4aa95def75f96.E_input', t33_wellord2)).
fof(c_0_17, plain, ![X5, X6]:r1_tarski(X5,k2_xboole_0(X5,X6)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_18, plain, ![X7]:(~v1_relat_1(X7)|k3_relat_1(X7)=k2_xboole_0(k1_relat_1(X7),k2_relat_1(X7))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_relat_1])])).
fof(c_0_19, plain, ![X8, X9, X10, X11]:(((k3_relat_1(X9)=X8|X9!=k1_wellord2(X8)|~v1_relat_1(X9))&((~r2_hidden(k4_tarski(X10,X11),X9)|r1_tarski(X10,X11)|(~r2_hidden(X10,X8)|~r2_hidden(X11,X8))|X9!=k1_wellord2(X8)|~v1_relat_1(X9))&(~r1_tarski(X10,X11)|r2_hidden(k4_tarski(X10,X11),X9)|(~r2_hidden(X10,X8)|~r2_hidden(X11,X8))|X9!=k1_wellord2(X8)|~v1_relat_1(X9))))&(((r2_hidden(esk1_2(X8,X9),X8)|k3_relat_1(X9)!=X8|X9=k1_wellord2(X8)|~v1_relat_1(X9))&(r2_hidden(esk2_2(X8,X9),X8)|k3_relat_1(X9)!=X8|X9=k1_wellord2(X8)|~v1_relat_1(X9)))&((~r2_hidden(k4_tarski(esk1_2(X8,X9),esk2_2(X8,X9)),X9)|~r1_tarski(esk1_2(X8,X9),esk2_2(X8,X9))|k3_relat_1(X9)!=X8|X9=k1_wellord2(X8)|~v1_relat_1(X9))&(r2_hidden(k4_tarski(esk1_2(X8,X9),esk2_2(X8,X9)),X9)|r1_tarski(esk1_2(X8,X9),esk2_2(X8,X9))|k3_relat_1(X9)!=X8|X9=k1_wellord2(X8)|~v1_relat_1(X9))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_wellord2])])])])])).
fof(c_0_20, plain, ![X14]:v1_relat_1(k1_wellord2(X14)), inference(variable_rename,[status(thm)],[dt_k1_wellord2])).
cnf(c_0_21, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, plain, (k3_relat_1(X1)=k2_xboole_0(k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, plain, (k3_relat_1(X1)=X2|X1!=k1_wellord2(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, plain, (v1_relat_1(k1_wellord2(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_25, plain, ![X15, X16]:(((r1_tarski(X15,X16)|X15!=X16)&(r1_tarski(X16,X15)|X15!=X16))&(~r1_tarski(X15,X16)|~r1_tarski(X16,X15)|X15=X16)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_26, plain, (r1_tarski(k1_relat_1(X1),k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27, plain, (k3_relat_1(k1_wellord2(X1))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_23]), c_0_24])])).
fof(c_0_28, plain, ![X17, X18]:((r1_tarski(k1_relat_1(X17),k1_relat_1(X18))|~r1_tarski(X17,X18)|~v1_relat_1(X18)|~v1_relat_1(X17))&(r1_tarski(k2_relat_1(X17),k2_relat_1(X18))|~r1_tarski(X17,X18)|~v1_relat_1(X18)|~v1_relat_1(X17))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t25_relat_1])])])])).
fof(c_0_29, plain, ![X19]:(k1_relat_1(k6_relat_1(X19))=X19&k2_relat_1(k6_relat_1(X19))=X19), inference(variable_rename,[status(thm)],[t71_relat_1])).
fof(c_0_30, plain, ![X20]:(v1_relat_1(k6_relat_1(X20))&v1_funct_1(k6_relat_1(X20))), inference(variable_rename,[status(thm)],[fc3_funct_1])).
fof(c_0_31, plain, ![X21, X22]:((r2_hidden(esk3_2(X21,X22),X21)|r1_tarski(k6_relat_1(X21),X22)|~v1_relat_1(X22))&(~r2_hidden(k4_tarski(esk3_2(X21,X22),esk3_2(X21,X22)),X22)|r1_tarski(k6_relat_1(X21),X22)|~v1_relat_1(X22))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t73_relat_1])])])])).
cnf(c_0_32, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r1_tarski(X1,X2)|~r2_hidden(X1,X4)|~r2_hidden(X2,X4)|X3!=k1_wellord2(X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_33, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_34, plain, ![X24, X25]:(~v1_relat_1(X25)|r1_tarski(k7_relat_1(X25,X24),X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t88_relat_1])])).
fof(c_0_35, plain, ![X26, X27]:(~v1_relat_1(X27)|k2_wellord1(X27,X26)=k7_relat_1(k8_relat_1(X26,X27),X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t17_wellord1])])).
fof(c_0_36, plain, ![X28, X29]:(~v1_relat_1(X29)|v1_relat_1(k8_relat_1(X28,X29))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_relat_1])])).
fof(c_0_37, plain, ![X30]:(~v1_relat_1(X30)|k2_wellord1(X30,k3_relat_1(X30))=X30), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_wellord1])])).
cnf(c_0_38, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_39, plain, (r1_tarski(k1_relat_1(k1_wellord2(X1)),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_24])])).
cnf(c_0_40, plain, (r1_tarski(k1_relat_1(X1),k1_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_41, plain, (k1_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_42, plain, (v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_43, plain, (r1_tarski(k6_relat_1(X1),X2)|~r2_hidden(k4_tarski(esk3_2(X1,X2),esk3_2(X1,X2)),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_44, plain, (r2_hidden(k4_tarski(X1,X2),k1_wellord2(X3))|~r2_hidden(X2,X3)|~r2_hidden(X1,X3)|~r1_tarski(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_32]), c_0_24])])).
cnf(c_0_45, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_33])).
fof(c_0_46, plain, ![X31, X32]:(~v1_relat_1(X32)|r1_tarski(k8_relat_1(X31,X32),X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t117_relat_1])])).
cnf(c_0_47, plain, (r1_tarski(k7_relat_1(X1,X2),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_48, plain, (k2_wellord1(X1,X2)=k7_relat_1(k8_relat_1(X2,X1),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_49, plain, (v1_relat_1(k8_relat_1(X2,X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_50, plain, (k2_wellord1(X1,k3_relat_1(X1))=X1|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_51, plain, (k1_relat_1(k1_wellord2(X1))=X1|~r1_tarski(X1,k1_relat_1(k1_wellord2(X1)))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_52, plain, (r1_tarski(X1,k1_relat_1(X2))|~v1_relat_1(X2)|~r1_tarski(k6_relat_1(X1),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42])])).
cnf(c_0_53, plain, (r1_tarski(k6_relat_1(X1),k1_wellord2(X2))|~r2_hidden(esk3_2(X1,k1_wellord2(X2)),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_24]), c_0_45])])).
cnf(c_0_54, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_tarski(k6_relat_1(X1),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_55, plain, ![X33, X34]:(~v1_relat_1(X34)|r1_tarski(k2_relat_1(k8_relat_1(X33,X34)),X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t116_relat_1])])).
cnf(c_0_56, plain, (r1_tarski(k8_relat_1(X2,X1),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_57, plain, (r1_tarski(k2_wellord1(X1,X2),k8_relat_1(X2,X1))|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])).
cnf(c_0_58, plain, (k2_wellord1(k1_wellord2(X1),X1)=k1_wellord2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_27]), c_0_24])])).
cnf(c_0_59, plain, (r1_tarski(k2_relat_1(X1),k2_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_60, plain, (k2_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_61, plain, ![X35]:(~v1_relat_1(X35)|r1_tarski(X35,k2_zfmisc_1(k1_relat_1(X35),k2_relat_1(X35)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t21_relat_1])])).
cnf(c_0_62, plain, (k1_relat_1(k1_wellord2(X1))=X1|~r1_tarski(k6_relat_1(X1),k1_wellord2(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_24])])).
cnf(c_0_63, plain, (r1_tarski(k6_relat_1(X1),k1_wellord2(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_24])])).
cnf(c_0_64, plain, (r1_tarski(k2_relat_1(k8_relat_1(X2,X1)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_65, plain, (k8_relat_1(X1,X2)=X2|~v1_relat_1(X2)|~r1_tarski(X2,k8_relat_1(X1,X2))), inference(spm,[status(thm)],[c_0_38, c_0_56])).
cnf(c_0_66, plain, (r1_tarski(k1_wellord2(X1),k8_relat_1(X1,k1_wellord2(X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_24])])).
cnf(c_0_67, plain, (r1_tarski(X1,k2_relat_1(X2))|~v1_relat_1(X2)|~r1_tarski(k6_relat_1(X1),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_42])])).
fof(c_0_68, negated_conjecture, ~(![X1]:r1_tarski(k1_wellord2(X1),k2_zfmisc_1(X1,X1))), inference(assume_negation,[status(cth)],[t33_wellord2])).
cnf(c_0_69, plain, (r1_tarski(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_70, plain, (k1_relat_1(k1_wellord2(X1))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_63])])).
cnf(c_0_71, plain, (k2_relat_1(k8_relat_1(X1,X2))=X1|~v1_relat_1(X2)|~r1_tarski(X1,k2_relat_1(k8_relat_1(X1,X2)))), inference(spm,[status(thm)],[c_0_38, c_0_64])).
cnf(c_0_72, plain, (k8_relat_1(X1,k1_wellord2(X1))=k1_wellord2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_24])])).
cnf(c_0_73, plain, (r1_tarski(X1,k2_relat_1(k1_wellord2(X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_63]), c_0_24])])).
fof(c_0_74, negated_conjecture, ~r1_tarski(k1_wellord2(esk4_0),k2_zfmisc_1(esk4_0,esk4_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_68])])])).
cnf(c_0_75, plain, (r1_tarski(k1_wellord2(X1),k2_zfmisc_1(X1,k2_relat_1(k1_wellord2(X1))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_24])])).
cnf(c_0_76, plain, (k2_relat_1(k1_wellord2(X1))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_24]), c_0_73])])).
cnf(c_0_77, negated_conjecture, (~r1_tarski(k1_wellord2(esk4_0),k2_zfmisc_1(esk4_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_78, plain, (r1_tarski(k1_wellord2(X1),k2_zfmisc_1(X1,X1))), inference(rw,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_79, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_78])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 80
# Proof object clause steps            : 45
# Proof object formula steps           : 35
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 39
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 29
# Processed clauses                    : 124
# ...of these trivial                  : 4
# ...subsumed                          : 10
# ...remaining for further processing  : 110
# Other redundant clauses eliminated   : 9
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 14
# Generated clauses                    : 140
# ...of the previous two non-trivial   : 121
# Contextual simplify-reflections      : 1
# Paramodulations                      : 131
# Factorizations                       : 0
# Equation resolutions                 : 9
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 59
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 42
# Current number of unprocessed clauses: 54
# ...number of literals in the above   : 174
# Current number of archived formulas  : 0
# Current number of archived clauses   : 42
# Clause-clause subsumption calls (NU) : 441
# Rec. Clause-clause subsumption calls : 328
# Non-unit clause-clause subsumptions  : 11
# Unit Clause-clause subsumption calls : 51
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 13
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4073

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28852 pages
