# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t38_wellord1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>((v2_wellord1(X3)&(r2_hidden(X1,k3_relat_1(X3))&r2_hidden(X2,k3_relat_1(X3))))=>(r1_tarski(k1_wellord1(X3,X1),k1_wellord1(X3,X2))<=>(X1=X2|r2_hidden(X1,k1_wellord1(X3,X2)))))), file('data/MPTP/proofs/5f273f1dc9d044b6bab51b43f7394689.E_input', t38_wellord1)).
fof(d1_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(X3=k1_wellord1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(~(X4=X2)&r2_hidden(k4_tarski(X4,X2),X1))))), file('data/MPTP/proofs/5f273f1dc9d044b6bab51b43f7394689.E_input', d1_wellord1)).
fof(t37_wellord1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>((v2_wellord1(X3)&(r2_hidden(X1,k3_relat_1(X3))&r2_hidden(X2,k3_relat_1(X3))))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>r1_tarski(k1_wellord1(X3,X1),k1_wellord1(X3,X2))))), file('data/MPTP/proofs/5f273f1dc9d044b6bab51b43f7394689.E_input', t37_wellord1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/5f273f1dc9d044b6bab51b43f7394689.E_input', d10_xboole_0)).
fof(c_0_4, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>((v2_wellord1(X3)&(r2_hidden(X1,k3_relat_1(X3))&r2_hidden(X2,k3_relat_1(X3))))=>(r1_tarski(k1_wellord1(X3,X1),k1_wellord1(X3,X2))<=>(X1=X2|r2_hidden(X1,k1_wellord1(X3,X2))))))), inference(assume_negation,[status(cth)],[t38_wellord1])).
fof(c_0_5, plain, ![X13, X14, X15, X16, X17, X18, X19]:((((X16!=X14|~r2_hidden(X16,X15)|X15!=k1_wellord1(X13,X14)|~v1_relat_1(X13))&(r2_hidden(k4_tarski(X16,X14),X13)|~r2_hidden(X16,X15)|X15!=k1_wellord1(X13,X14)|~v1_relat_1(X13)))&(X17=X14|~r2_hidden(k4_tarski(X17,X14),X13)|r2_hidden(X17,X15)|X15!=k1_wellord1(X13,X14)|~v1_relat_1(X13)))&((~r2_hidden(esk4_3(X13,X18,X19),X19)|(esk4_3(X13,X18,X19)=X18|~r2_hidden(k4_tarski(esk4_3(X13,X18,X19),X18),X13))|X19=k1_wellord1(X13,X18)|~v1_relat_1(X13))&((esk4_3(X13,X18,X19)!=X18|r2_hidden(esk4_3(X13,X18,X19),X19)|X19=k1_wellord1(X13,X18)|~v1_relat_1(X13))&(r2_hidden(k4_tarski(esk4_3(X13,X18,X19),X18),X13)|r2_hidden(esk4_3(X13,X18,X19),X19)|X19=k1_wellord1(X13,X18)|~v1_relat_1(X13))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_wellord1])])])])])])])).
fof(c_0_6, plain, ![X10, X11, X12]:((~r2_hidden(k4_tarski(X10,X11),X12)|r1_tarski(k1_wellord1(X12,X10),k1_wellord1(X12,X11))|(~v2_wellord1(X12)|(~r2_hidden(X10,k3_relat_1(X12))|~r2_hidden(X11,k3_relat_1(X12))))|~v1_relat_1(X12))&(~r1_tarski(k1_wellord1(X12,X10),k1_wellord1(X12,X11))|r2_hidden(k4_tarski(X10,X11),X12)|(~v2_wellord1(X12)|(~r2_hidden(X10,k3_relat_1(X12))|~r2_hidden(X11,k3_relat_1(X12))))|~v1_relat_1(X12))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t37_wellord1])])])).
fof(c_0_7, negated_conjecture, (v1_relat_1(esk3_0)&((v2_wellord1(esk3_0)&(r2_hidden(esk1_0,k3_relat_1(esk3_0))&r2_hidden(esk2_0,k3_relat_1(esk3_0))))&(((esk1_0!=esk2_0|~r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0)))&(~r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))|~r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))))&(r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))|(esk1_0=esk2_0|r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
cnf(c_0_8, plain, (X1=X2|r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_wellord1(X3,X2)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (r2_hidden(k4_tarski(X2,X3),X1)|~r1_tarski(k1_wellord1(X1,X2),k1_wellord1(X1,X3))|~v2_wellord1(X1)|~r2_hidden(X2,k3_relat_1(X1))|~r2_hidden(X3,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, negated_conjecture, (r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))|esk1_0=esk2_0|r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, negated_conjecture, (r2_hidden(esk2_0,k3_relat_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, negated_conjecture, (r2_hidden(esk1_0,k3_relat_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, negated_conjecture, (v2_wellord1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_15, negated_conjecture, (~r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))|~r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_16, plain, (r1_tarski(k1_wellord1(X3,X1),k1_wellord1(X3,X2))|~r2_hidden(k4_tarski(X1,X2),X3)|~v2_wellord1(X3)|~r2_hidden(X1,k3_relat_1(X3))|~r2_hidden(X2,k3_relat_1(X3))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_17, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,X4)|X4!=k1_wellord1(X3,X2)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_18, plain, (X1=X2|r2_hidden(X1,k1_wellord1(X3,X2))|~r2_hidden(k4_tarski(X1,X2),X3)|~v1_relat_1(X3)), inference(er,[status(thm)],[c_0_8])).
cnf(c_0_19, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11]), c_0_12]), c_0_13]), c_0_14])])).
cnf(c_0_20, negated_conjecture, (~r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))|~r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_11]), c_0_12]), c_0_13]), c_0_14])])).
cnf(c_0_21, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,k1_wellord1(X3,X2))|~v1_relat_1(X3)), inference(er,[status(thm)],[c_0_17])).
fof(c_0_22, plain, ![X8, X9]:(((r1_tarski(X8,X9)|X8!=X9)&(r1_tarski(X9,X8)|X8!=X9))&(~r1_tarski(X8,X9)|~r1_tarski(X9,X8)|X8=X9)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_23, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_14])])).
cnf(c_0_24, negated_conjecture, (~r2_hidden(esk1_0,k1_wellord1(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_14])])).
cnf(c_0_25, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, negated_conjecture, (esk1_0!=esk2_0|~r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_27, negated_conjecture, (esk2_0=esk1_0), inference(sr,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_28, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_25])).
cnf(c_0_29, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_27]), c_0_27]), c_0_28])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 30
# Proof object clause steps            : 21
# Proof object formula steps           : 9
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 4
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 18
# Processed clauses                    : 48
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 48
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 7
# Generated clauses                    : 24
# ...of the previous two non-trivial   : 23
# Contextual simplify-reflections      : 0
# Paramodulations                      : 18
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 15
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 10
# Current number of unprocessed clauses: 5
# ...number of literals in the above   : 23
# Current number of archived formulas  : 0
# Current number of archived clauses   : 28
# Clause-clause subsumption calls (NU) : 84
# Rec. Clause-clause subsumption calls : 30
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 3
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1683

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.000 s
# Total time               : 0.009 s
# Maximum resident set size: 28656 pages
