# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d8_filter_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v3_filter_0(X1)&l3_lattices(X1))))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k4_filter_0(X1,X2,X3)<=>(r3_lattices(X1,k4_lattices(X1,X2,X4),X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r3_lattices(X1,k4_lattices(X1,X2,X5),X3)=>r3_lattices(X1,X5,X4)))))))))), file('data/MPTP/proofs/5f7af09cb850426395ca74125850cf8d.E_input', d8_filter_0)).
fof(fraenkel_a_3_6_lattice3, axiom, ![X1, X2, X3, X4]:(((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))&(m1_subset_1(X3,u1_struct_0(X2))&m1_subset_1(X4,u1_struct_0(X2))))=>(r2_hidden(X1,a_3_6_lattice3(X2,X3,X4))<=>?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(X1=X5&r3_lattices(X2,k4_lattices(X2,X3,X5),X4))))), file('data/MPTP/proofs/5f7af09cb850426395ca74125850cf8d.E_input', fraenkel_a_3_6_lattice3)).
fof(d21_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X3=k15_lattice3(X1,X2)<=>(r4_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r4_lattice3(X1,X4,X2)=>r1_lattices(X1,X3,X4)))))))), file('data/MPTP/proofs/5f7af09cb850426395ca74125850cf8d.E_input', d21_lattice3)).
fof(redefinition_r3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_lattices(X1,X2,X3)<=>r1_lattices(X1,X2,X3))), file('data/MPTP/proofs/5f7af09cb850426395ca74125850cf8d.E_input', redefinition_r3_lattices)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/5f7af09cb850426395ca74125850cf8d.E_input', cc1_lattices)).
fof(d17_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r4_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X4,X2)))))), file('data/MPTP/proofs/5f7af09cb850426395ca74125850cf8d.E_input', d17_lattice3)).
fof(dt_k4_filter_0, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k4_filter_0(X1,X2,X3),u1_struct_0(X1))), file('data/MPTP/proofs/5f7af09cb850426395ca74125850cf8d.E_input', dt_k4_filter_0)).
fof(t52_lattice3, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v3_filter_0(X1)&l3_lattices(X1))))=>k4_filter_0(X1,X2,X3)=k15_lattice3(X1,a_3_6_lattice3(X1,X2,X3)))))), file('data/MPTP/proofs/5f7af09cb850426395ca74125850cf8d.E_input', t52_lattice3)).
fof(c_0_8, plain, ![X6, X7, X8, X9, X10]:(((r3_lattices(X6,k4_lattices(X6,X7,X9),X8)|X9!=k4_filter_0(X6,X7,X8)|~m1_subset_1(X9,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v3_filter_0(X6)|~l3_lattices(X6))))|~m1_subset_1(X8,u1_struct_0(X6))|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|~l3_lattices(X6))))&(~m1_subset_1(X10,u1_struct_0(X6))|(~r3_lattices(X6,k4_lattices(X6,X7,X10),X8)|r3_lattices(X6,X10,X9))|X9!=k4_filter_0(X6,X7,X8)|~m1_subset_1(X9,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v3_filter_0(X6)|~l3_lattices(X6))))|~m1_subset_1(X8,u1_struct_0(X6))|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|~l3_lattices(X6)))))&((m1_subset_1(esk1_4(X6,X7,X8,X9),u1_struct_0(X6))|~r3_lattices(X6,k4_lattices(X6,X7,X9),X8)|X9=k4_filter_0(X6,X7,X8)|~m1_subset_1(X9,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v3_filter_0(X6)|~l3_lattices(X6))))|~m1_subset_1(X8,u1_struct_0(X6))|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|~l3_lattices(X6))))&((r3_lattices(X6,k4_lattices(X6,X7,esk1_4(X6,X7,X8,X9)),X8)|~r3_lattices(X6,k4_lattices(X6,X7,X9),X8)|X9=k4_filter_0(X6,X7,X8)|~m1_subset_1(X9,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v3_filter_0(X6)|~l3_lattices(X6))))|~m1_subset_1(X8,u1_struct_0(X6))|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|~l3_lattices(X6))))&(~r3_lattices(X6,esk1_4(X6,X7,X8,X9),X9)|~r3_lattices(X6,k4_lattices(X6,X7,X9),X8)|X9=k4_filter_0(X6,X7,X8)|~m1_subset_1(X9,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v3_filter_0(X6)|~l3_lattices(X6))))|~m1_subset_1(X8,u1_struct_0(X6))|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|~l3_lattices(X6))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_filter_0])])])])])])).
cnf(c_0_9, plain, (r3_lattices(X2,X1,X5)|v2_struct_0(X2)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r3_lattices(X2,k4_lattices(X2,X3,X1),X4)|X5!=k4_filter_0(X2,X3,X4)|~m1_subset_1(X5,u1_struct_0(X2))|~v10_lattices(X2)|~v3_filter_0(X2)|~l3_lattices(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~v10_lattices(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_10, plain, ![X12, X13, X14, X15, X17]:(((m1_subset_1(esk2_4(X12,X13,X14,X15),u1_struct_0(X13))|~r2_hidden(X12,a_3_6_lattice3(X13,X14,X15))|(v2_struct_0(X13)|(~v10_lattices(X13)|(~v4_lattice3(X13)|~l3_lattices(X13)))|(~m1_subset_1(X14,u1_struct_0(X13))|~m1_subset_1(X15,u1_struct_0(X13)))))&((X12=esk2_4(X12,X13,X14,X15)|~r2_hidden(X12,a_3_6_lattice3(X13,X14,X15))|(v2_struct_0(X13)|(~v10_lattices(X13)|(~v4_lattice3(X13)|~l3_lattices(X13)))|(~m1_subset_1(X14,u1_struct_0(X13))|~m1_subset_1(X15,u1_struct_0(X13)))))&(r3_lattices(X13,k4_lattices(X13,X14,esk2_4(X12,X13,X14,X15)),X15)|~r2_hidden(X12,a_3_6_lattice3(X13,X14,X15))|(v2_struct_0(X13)|(~v10_lattices(X13)|(~v4_lattice3(X13)|~l3_lattices(X13)))|(~m1_subset_1(X14,u1_struct_0(X13))|~m1_subset_1(X15,u1_struct_0(X13)))))))&(~m1_subset_1(X17,u1_struct_0(X13))|(X12!=X17|~r3_lattices(X13,k4_lattices(X13,X14,X17),X15))|r2_hidden(X12,a_3_6_lattice3(X13,X14,X15))|(v2_struct_0(X13)|(~v10_lattices(X13)|(~v4_lattice3(X13)|~l3_lattices(X13)))|(~m1_subset_1(X14,u1_struct_0(X13))|~m1_subset_1(X15,u1_struct_0(X13)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_3_6_lattice3])])])])])])).
fof(c_0_11, plain, ![X18, X19, X20, X21]:(((r4_lattice3(X18,X20,X19)|X20!=k15_lattice3(X18,X19)|~m1_subset_1(X20,u1_struct_0(X18))|(v2_struct_0(X18)|(~v10_lattices(X18)|(~v4_lattice3(X18)|~l3_lattices(X18))))|(v2_struct_0(X18)|~l3_lattices(X18)))&(~m1_subset_1(X21,u1_struct_0(X18))|(~r4_lattice3(X18,X21,X19)|r1_lattices(X18,X20,X21))|X20!=k15_lattice3(X18,X19)|~m1_subset_1(X20,u1_struct_0(X18))|(v2_struct_0(X18)|(~v10_lattices(X18)|(~v4_lattice3(X18)|~l3_lattices(X18))))|(v2_struct_0(X18)|~l3_lattices(X18))))&((m1_subset_1(esk3_3(X18,X19,X20),u1_struct_0(X18))|~r4_lattice3(X18,X20,X19)|X20=k15_lattice3(X18,X19)|~m1_subset_1(X20,u1_struct_0(X18))|(v2_struct_0(X18)|(~v10_lattices(X18)|(~v4_lattice3(X18)|~l3_lattices(X18))))|(v2_struct_0(X18)|~l3_lattices(X18)))&((r4_lattice3(X18,esk3_3(X18,X19,X20),X19)|~r4_lattice3(X18,X20,X19)|X20=k15_lattice3(X18,X19)|~m1_subset_1(X20,u1_struct_0(X18))|(v2_struct_0(X18)|(~v10_lattices(X18)|(~v4_lattice3(X18)|~l3_lattices(X18))))|(v2_struct_0(X18)|~l3_lattices(X18)))&(~r1_lattices(X18,X20,esk3_3(X18,X19,X20))|~r4_lattice3(X18,X20,X19)|X20=k15_lattice3(X18,X19)|~m1_subset_1(X20,u1_struct_0(X18))|(v2_struct_0(X18)|(~v10_lattices(X18)|(~v4_lattice3(X18)|~l3_lattices(X18))))|(v2_struct_0(X18)|~l3_lattices(X18)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d21_lattice3])])])])])])).
fof(c_0_12, plain, ![X23, X24, X25]:((~r3_lattices(X23,X24,X25)|r1_lattices(X23,X24,X25)|(v2_struct_0(X23)|(~v6_lattices(X23)|(~v8_lattices(X23)|(~v9_lattices(X23)|~l3_lattices(X23))))|(~m1_subset_1(X24,u1_struct_0(X23))|~m1_subset_1(X25,u1_struct_0(X23)))))&(~r1_lattices(X23,X24,X25)|r3_lattices(X23,X24,X25)|(v2_struct_0(X23)|(~v6_lattices(X23)|(~v8_lattices(X23)|(~v9_lattices(X23)|~l3_lattices(X23))))|(~m1_subset_1(X24,u1_struct_0(X23))|~m1_subset_1(X25,u1_struct_0(X23)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_lattices])])])])).
fof(c_0_13, plain, ![X26]:((~v2_struct_0(X26)|(v2_struct_0(X26)|~v10_lattices(X26))|~l3_lattices(X26))&((v4_lattices(X26)|(v2_struct_0(X26)|~v10_lattices(X26))|~l3_lattices(X26))&((v5_lattices(X26)|(v2_struct_0(X26)|~v10_lattices(X26))|~l3_lattices(X26))&((v6_lattices(X26)|(v2_struct_0(X26)|~v10_lattices(X26))|~l3_lattices(X26))&((v7_lattices(X26)|(v2_struct_0(X26)|~v10_lattices(X26))|~l3_lattices(X26))&((v8_lattices(X26)|(v2_struct_0(X26)|~v10_lattices(X26))|~l3_lattices(X26))&(v9_lattices(X26)|(v2_struct_0(X26)|~v10_lattices(X26))|~l3_lattices(X26)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
cnf(c_0_14, plain, (v2_struct_0(X2)|r3_lattices(X2,X1,X5)|X5!=k4_filter_0(X2,X3,X4)|~v10_lattices(X2)|~l3_lattices(X2)|~v3_filter_0(X2)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~r3_lattices(X2,k4_lattices(X2,X3,X1),X4)), inference(cn,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (r3_lattices(X1,k4_lattices(X1,X2,esk2_4(X3,X1,X2,X4)),X4)|v2_struct_0(X1)|~r2_hidden(X3,a_3_6_lattice3(X1,X2,X4))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (m1_subset_1(esk2_4(X1,X2,X3,X4),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_3_6_lattice3(X2,X3,X4))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_17, plain, ![X27, X28, X29, X30, X31]:((~r4_lattice3(X27,X28,X29)|(~m1_subset_1(X30,u1_struct_0(X27))|(~r2_hidden(X30,X29)|r1_lattices(X27,X30,X28)))|~m1_subset_1(X28,u1_struct_0(X27))|(v2_struct_0(X27)|~l3_lattices(X27)))&((m1_subset_1(esk4_3(X27,X28,X31),u1_struct_0(X27))|r4_lattice3(X27,X28,X31)|~m1_subset_1(X28,u1_struct_0(X27))|(v2_struct_0(X27)|~l3_lattices(X27)))&((r2_hidden(esk4_3(X27,X28,X31),X31)|r4_lattice3(X27,X28,X31)|~m1_subset_1(X28,u1_struct_0(X27))|(v2_struct_0(X27)|~l3_lattices(X27)))&(~r1_lattices(X27,esk4_3(X27,X28,X31),X28)|r4_lattice3(X27,X28,X31)|~m1_subset_1(X28,u1_struct_0(X27))|(v2_struct_0(X27)|~l3_lattices(X27)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_lattice3])])])])])])])).
cnf(c_0_18, plain, (r4_lattice3(X1,esk3_3(X1,X2,X3),X2)|X3=k15_lattice3(X1,X2)|v2_struct_0(X1)|v2_struct_0(X1)|~r4_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19, plain, (m1_subset_1(esk3_3(X1,X2,X3),u1_struct_0(X1))|X3=k15_lattice3(X1,X2)|v2_struct_0(X1)|v2_struct_0(X1)|~r4_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20, plain, (r1_lattices(X1,X2,X3)|v2_struct_0(X1)|~r3_lattices(X1,X2,X3)|~v6_lattices(X1)|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, plain, (v9_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22, plain, (v6_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23, plain, (v8_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_24, plain, (r3_lattices(X1,esk2_4(X2,X1,X3,X4),X5)|v2_struct_0(X1)|X5!=k4_filter_0(X1,X3,X4)|~r2_hidden(X2,a_3_6_lattice3(X1,X3,X4))|~v4_lattice3(X1)|~v3_filter_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])).
cnf(c_0_25, plain, (X1=esk2_4(X1,X2,X3,X4)|v2_struct_0(X2)|~r2_hidden(X1,a_3_6_lattice3(X2,X3,X4))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_26, plain, (X2=k15_lattice3(X1,X3)|v2_struct_0(X1)|v2_struct_0(X1)|~r1_lattices(X1,X2,esk3_3(X1,X3,X2))|~r4_lattice3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_27, plain, (r1_lattices(X1,X4,X2)|v2_struct_0(X1)|~r4_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28, plain, (X3=k15_lattice3(X1,X2)|v2_struct_0(X1)|r4_lattice3(X1,esk3_3(X1,X2,X3),X2)|~v10_lattices(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~r4_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_18])).
cnf(c_0_29, plain, (X3=k15_lattice3(X1,X2)|v2_struct_0(X1)|m1_subset_1(esk3_3(X1,X2,X3),u1_struct_0(X1))|~v10_lattices(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~r4_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_19])).
cnf(c_0_30, plain, (r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r1_lattices(X1,esk4_3(X1,X2,X3),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_31, plain, (r1_lattices(X1,X2,X3)|v2_struct_0(X1)|~r3_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23])).
cnf(c_0_32, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X1))|r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_33, plain, (r3_lattices(X1,X2,X3)|v2_struct_0(X1)|X3!=k4_filter_0(X1,X4,X5)|~r2_hidden(X2,a_3_6_lattice3(X1,X4,X5))|~v4_lattice3(X1)|~v3_filter_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_34, plain, (r2_hidden(esk4_3(X1,X2,X3),X3)|r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_35, plain, (X2=k15_lattice3(X1,X3)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~r4_lattice3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattices(X1,X2,esk3_3(X1,X3,X2))), inference(cn,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (X1=k15_lattice3(X2,X3)|r1_lattices(X2,X4,esk3_3(X2,X3,X1))|v2_struct_0(X2)|~r4_lattice3(X2,X1,X3)|~r2_hidden(X4,X3)|~v4_lattice3(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v10_lattices(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])).
cnf(c_0_37, plain, (r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r3_lattices(X1,esk4_3(X1,X2,X3),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])).
cnf(c_0_38, plain, (r4_lattice3(X1,X2,a_3_6_lattice3(X3,X4,X5))|r3_lattices(X3,esk4_3(X1,X2,a_3_6_lattice3(X3,X4,X5)),X6)|v2_struct_0(X1)|v2_struct_0(X3)|X6!=k4_filter_0(X3,X4,X5)|~v4_lattice3(X3)|~v3_filter_0(X3)|~m1_subset_1(X6,u1_struct_0(X3))|~m1_subset_1(X5,u1_struct_0(X3))|~m1_subset_1(X4,u1_struct_0(X3))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X3)|~l3_lattices(X1)|~v10_lattices(X3)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_39, plain, (X1=k15_lattice3(X2,X3)|v2_struct_0(X2)|~r4_lattice3(X2,X1,X3)|~r2_hidden(X1,X3)|~v4_lattice3(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v10_lattices(X2)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_40, plain, (r4_lattice3(X1,X2,a_3_6_lattice3(X1,X3,X4))|v2_struct_0(X1)|X2!=k4_filter_0(X1,X3,X4)|~v4_lattice3(X1)|~v3_filter_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_41, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_3_6_lattice3(X2,X3,X4))|~v4_lattice3(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_lattices(X2)|~v10_lattices(X2)), inference(spm,[status(thm)],[c_0_16, c_0_25])).
cnf(c_0_42, plain, (r2_hidden(X3,a_3_6_lattice3(X2,X4,X5))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~r3_lattices(X2,k4_lattices(X2,X4,X1),X5)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_43, plain, (r3_lattices(X1,k4_lattices(X1,X2,X3),X4)|v2_struct_0(X1)|v2_struct_0(X1)|X3!=k4_filter_0(X1,X2,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~v10_lattices(X1)|~v3_filter_0(X1)|~l3_lattices(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_44, plain, (X1=k15_lattice3(X2,a_3_6_lattice3(X2,X3,X4))|v2_struct_0(X2)|X1!=k4_filter_0(X2,X3,X4)|~r2_hidden(X1,a_3_6_lattice3(X2,X3,X4))|~v4_lattice3(X2)|~v3_filter_0(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_lattices(X2)|~v10_lattices(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])).
cnf(c_0_45, plain, (r2_hidden(X1,a_3_6_lattice3(X2,X3,X4))|v2_struct_0(X2)|~v4_lattice3(X2)|~r3_lattices(X2,k4_lattices(X2,X3,X1),X4)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v10_lattices(X2)), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_46, plain, (v2_struct_0(X1)|r3_lattices(X1,k4_lattices(X1,X2,X3),X4)|X3!=k4_filter_0(X1,X2,X4)|~v10_lattices(X1)|~l3_lattices(X1)|~v3_filter_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_43])).
fof(c_0_47, plain, ![X33, X34, X35]:(v2_struct_0(X33)|(~v10_lattices(X33)|~l3_lattices(X33))|(~m1_subset_1(X34,u1_struct_0(X33))|~m1_subset_1(X35,u1_struct_0(X33)))|m1_subset_1(k4_filter_0(X33,X34,X35),u1_struct_0(X33))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_filter_0])])])).
fof(c_0_48, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v3_filter_0(X1)&l3_lattices(X1))))=>k4_filter_0(X1,X2,X3)=k15_lattice3(X1,a_3_6_lattice3(X1,X2,X3))))))), inference(assume_negation,[status(cth)],[t52_lattice3])).
cnf(c_0_49, plain, (X1=k15_lattice3(X2,a_3_6_lattice3(X2,X3,X4))|v2_struct_0(X2)|X1!=k4_filter_0(X2,X3,X4)|~v4_lattice3(X2)|~v3_filter_0(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v10_lattices(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])).
cnf(c_0_50, plain, (v2_struct_0(X1)|m1_subset_1(k4_filter_0(X1,X2,X3),u1_struct_0(X1))|~v10_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_51, negated_conjecture, ((~v2_struct_0(esk5_0)&(v10_lattices(esk5_0)&(v4_lattice3(esk5_0)&l3_lattices(esk5_0))))&(m1_subset_1(esk6_0,u1_struct_0(esk5_0))&(m1_subset_1(esk7_0,u1_struct_0(esk5_0))&((~v2_struct_0(esk5_0)&(v10_lattices(esk5_0)&(v3_filter_0(esk5_0)&l3_lattices(esk5_0))))&k4_filter_0(esk5_0,esk6_0,esk7_0)!=k15_lattice3(esk5_0,a_3_6_lattice3(esk5_0,esk6_0,esk7_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_48])])])])).
cnf(c_0_52, plain, (k4_filter_0(X1,X2,X3)=k15_lattice3(X1,a_3_6_lattice3(X1,X4,X5))|v2_struct_0(X1)|k4_filter_0(X1,X2,X3)!=k4_filter_0(X1,X4,X5)|~v4_lattice3(X1)|~v3_filter_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_53, negated_conjecture, (k4_filter_0(esk5_0,esk6_0,esk7_0)!=k15_lattice3(esk5_0,a_3_6_lattice3(esk5_0,esk6_0,esk7_0))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_54, plain, (k15_lattice3(X1,a_3_6_lattice3(X1,X2,X3))=k4_filter_0(X1,X2,X3)|v2_struct_0(X1)|~v4_lattice3(X1)|~v3_filter_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(er,[status(thm)],[c_0_52])).
cnf(c_0_55, negated_conjecture, (v4_lattice3(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_56, negated_conjecture, (v3_filter_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_57, negated_conjecture, (m1_subset_1(esk7_0,u1_struct_0(esk5_0))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_58, negated_conjecture, (m1_subset_1(esk6_0,u1_struct_0(esk5_0))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_59, negated_conjecture, (l3_lattices(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_60, negated_conjecture, (v10_lattices(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_61, negated_conjecture, (~v2_struct_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_62, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55]), c_0_56]), c_0_57]), c_0_58]), c_0_59]), c_0_60])]), c_0_61]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 63
# Proof object clause steps            : 46
# Proof object formula steps           : 17
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 21
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 38
# Processed clauses                    : 128
# ...of these trivial                  : 2
# ...subsumed                          : 8
# ...remaining for further processing  : 118
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 107
# ...of the previous two non-trivial   : 88
# Contextual simplify-reflections      : 31
# Paramodulations                      : 102
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 81
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 73
# Current number of unprocessed clauses: 33
# ...number of literals in the above   : 479
# Current number of archived formulas  : 0
# Current number of archived clauses   : 36
# Clause-clause subsumption calls (NU) : 6867
# Rec. Clause-clause subsumption calls : 152
# Non-unit clause-clause subsumptions  : 39
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 10127

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 28792 pages
