# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t35_waybel_9, conjecture, ![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((~(v2_struct_0(X3))&(v4_orders_2(X3)&(v7_waybel_0(X3)&l1_waybel_0(X3,X1))))=>((![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X4),X1,X1))&(v10_waybel_0(X3,X1)&r3_waybel_9(X1,X3,X2)))=>X2=k1_waybel_2(X1,X3))))), file('data/MPTP/proofs/5f7c177f34ec4345b6af7cfbe9c2e0a3.E_input', t35_waybel_9)).
fof(cc2_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v2_lattice3(X1)=>~(v2_struct_0(X1)))), file('data/MPTP/proofs/5f7c177f34ec4345b6af7cfbe9c2e0a3.E_input', cc2_lattice3)).
fof(dt_l1_waybel_9, axiom, ![X1]:(l1_waybel_9(X1)=>(l1_pre_topc(X1)&l1_orders_2(X1))), file('data/MPTP/proofs/5f7c177f34ec4345b6af7cfbe9c2e0a3.E_input', dt_l1_waybel_9)).
fof(l48_waybel_9, axiom, ![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((X3=X4&(![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X5),X1,X1))&(v10_waybel_0(X2,X1)&r3_waybel_9(X1,X2,X3))))=>r2_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X4)))))), file('data/MPTP/proofs/5f7c177f34ec4345b6af7cfbe9c2e0a3.E_input', l48_waybel_9)).
fof(d1_waybel_2, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(l1_waybel_0(X2,X1)=>k1_waybel_2(X1,X2)=k4_yellow_2(X1,u1_waybel_0(X1,X2)))), file('data/MPTP/proofs/5f7c177f34ec4345b6af7cfbe9c2e0a3.E_input', d1_waybel_2)).
fof(dt_u1_waybel_0, axiom, ![X1, X2]:((l1_struct_0(X1)&l1_waybel_0(X2,X1))=>(v1_funct_1(u1_waybel_0(X1,X2))&(v1_funct_2(u1_waybel_0(X1,X2),u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(u1_waybel_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))), file('data/MPTP/proofs/5f7c177f34ec4345b6af7cfbe9c2e0a3.E_input', dt_u1_waybel_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/5f7c177f34ec4345b6af7cfbe9c2e0a3.E_input', dt_l1_orders_2)).
fof(d5_yellow_2, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(v1_relat_1(X2)=>k4_yellow_2(X1,X2)=k1_yellow_0(X1,k2_relat_1(X2)))), file('data/MPTP/proofs/5f7c177f34ec4345b6af7cfbe9c2e0a3.E_input', d5_yellow_2)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/5f7c177f34ec4345b6af7cfbe9c2e0a3.E_input', cc1_relset_1)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/5f7c177f34ec4345b6af7cfbe9c2e0a3.E_input', redefinition_k2_relset_1)).
fof(l49_waybel_9, axiom, ![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((X3=X4&(![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X5),X1,X1))&r3_waybel_9(X1,X2,X3)))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r2_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X5)=>r3_orders_2(X1,X4,X5)))))))), file('data/MPTP/proofs/5f7c177f34ec4345b6af7cfbe9c2e0a3.E_input', l49_waybel_9)).
fof(t30_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4)))))&((r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))=>(X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3)))))), file('data/MPTP/proofs/5f7c177f34ec4345b6af7cfbe9c2e0a3.E_input', t30_yellow_0)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/5f7c177f34ec4345b6af7cfbe9c2e0a3.E_input', redefinition_r3_orders_2)).
fof(c_0_13, negated_conjecture, ~(![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((~(v2_struct_0(X3))&(v4_orders_2(X3)&(v7_waybel_0(X3)&l1_waybel_0(X3,X1))))=>((![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X4),X1,X1))&(v10_waybel_0(X3,X1)&r3_waybel_9(X1,X3,X2)))=>X2=k1_waybel_2(X1,X3)))))), inference(assume_negation,[status(cth)],[t35_waybel_9])).
fof(c_0_14, plain, ![X10]:(~l1_orders_2(X10)|(~v2_lattice3(X10)|~v2_struct_0(X10))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_lattice3])])])).
fof(c_0_15, negated_conjecture, ![X9]:((v2_pre_topc(esk1_0)&(v8_pre_topc(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_lattice3(esk1_0)&(v2_lattice3(esk1_0)&(v1_compts_1(esk1_0)&l1_waybel_9(esk1_0)))))))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&((~v2_struct_0(esk3_0)&(v4_orders_2(esk3_0)&(v7_waybel_0(esk3_0)&l1_waybel_0(esk3_0,esk1_0))))&(((~m1_subset_1(X9,u1_struct_0(esk1_0))|v5_pre_topc(k4_waybel_1(esk1_0,X9),esk1_0,esk1_0))&(v10_waybel_0(esk3_0,esk1_0)&r3_waybel_9(esk1_0,esk3_0,esk2_0)))&esk2_0!=k1_waybel_2(esk1_0,esk3_0))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_13])])])])])).
fof(c_0_16, plain, ![X11]:((l1_pre_topc(X11)|~l1_waybel_9(X11))&(l1_orders_2(X11)|~l1_waybel_9(X11))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_waybel_9])])])).
cnf(c_0_17, plain, (~l1_orders_2(X1)|~v2_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, negated_conjecture, (v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, plain, (l1_orders_2(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, negated_conjecture, (l1_waybel_9(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_21, plain, ![X12, X13, X14, X15]:((m1_subset_1(esk4_4(X12,X13,X14,X15),u1_struct_0(X12))|(~v10_waybel_0(X13,X12)|~r3_waybel_9(X12,X13,X14))|X14!=X15|r2_lattice3(X12,k2_relset_1(u1_struct_0(X13),u1_struct_0(X12),u1_waybel_0(X12,X13)),X15)|~m1_subset_1(X15,u1_struct_0(X12))|~m1_subset_1(X14,u1_struct_0(X12))|(v2_struct_0(X13)|(~v4_orders_2(X13)|(~v7_waybel_0(X13)|~l1_waybel_0(X13,X12))))|(~v2_pre_topc(X12)|(~v8_pre_topc(X12)|(~v3_orders_2(X12)|(~v4_orders_2(X12)|(~v5_orders_2(X12)|(~v1_lattice3(X12)|(~v2_lattice3(X12)|(~v1_compts_1(X12)|~l1_waybel_9(X12))))))))))&(~v5_pre_topc(k4_waybel_1(X12,esk4_4(X12,X13,X14,X15)),X12,X12)|(~v10_waybel_0(X13,X12)|~r3_waybel_9(X12,X13,X14))|X14!=X15|r2_lattice3(X12,k2_relset_1(u1_struct_0(X13),u1_struct_0(X12),u1_waybel_0(X12,X13)),X15)|~m1_subset_1(X15,u1_struct_0(X12))|~m1_subset_1(X14,u1_struct_0(X12))|(v2_struct_0(X13)|(~v4_orders_2(X13)|(~v7_waybel_0(X13)|~l1_waybel_0(X13,X12))))|(~v2_pre_topc(X12)|(~v8_pre_topc(X12)|(~v3_orders_2(X12)|(~v4_orders_2(X12)|(~v5_orders_2(X12)|(~v1_lattice3(X12)|(~v2_lattice3(X12)|(~v1_compts_1(X12)|~l1_waybel_9(X12))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l48_waybel_9])])])])])])).
fof(c_0_22, plain, ![X17, X18]:(v2_struct_0(X17)|~l1_orders_2(X17)|(~l1_waybel_0(X18,X17)|k1_waybel_2(X17,X18)=k4_yellow_2(X17,u1_waybel_0(X17,X18)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_waybel_2])])])])).
cnf(c_0_23, negated_conjecture, (~l1_orders_2(esk1_0)|~v2_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_24, negated_conjecture, (l1_orders_2(esk1_0)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
fof(c_0_25, plain, ![X19, X20]:((v1_funct_1(u1_waybel_0(X19,X20))|(~l1_struct_0(X19)|~l1_waybel_0(X20,X19)))&((v1_funct_2(u1_waybel_0(X19,X20),u1_struct_0(X20),u1_struct_0(X19))|(~l1_struct_0(X19)|~l1_waybel_0(X20,X19)))&(m1_subset_1(u1_waybel_0(X19,X20),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X20),u1_struct_0(X19))))|(~l1_struct_0(X19)|~l1_waybel_0(X20,X19))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_waybel_0])])])).
cnf(c_0_26, plain, (m1_subset_1(esk4_4(X1,X2,X3,X4),u1_struct_0(X1))|r2_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X4)|v2_struct_0(X2)|~v10_waybel_0(X2,X1)|~r3_waybel_9(X1,X2,X3)|X3!=X4|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v1_compts_1(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (v2_struct_0(X1)|k1_waybel_2(X1,X2)=k4_yellow_2(X1,u1_waybel_0(X1,X2))|~l1_orders_2(X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, negated_conjecture, (l1_waybel_0(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_29, negated_conjecture, (~v2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_24])])).
cnf(c_0_30, plain, (m1_subset_1(u1_waybel_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_31, plain, ![X21]:(~l1_orders_2(X21)|l1_struct_0(X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_32, plain, (r2_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X3)|v2_struct_0(X2)|m1_subset_1(esk4_4(X1,X2,X3,X3),u1_struct_0(X1))|~r3_waybel_9(X1,X2,X3)|~v10_waybel_0(X2,X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_26])).
cnf(c_0_33, negated_conjecture, (v10_waybel_0(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_34, negated_conjecture, (v7_waybel_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_35, negated_conjecture, (v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_36, negated_conjecture, (v1_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_37, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_38, negated_conjecture, (v4_orders_2(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_39, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_40, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_41, negated_conjecture, (v8_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_42, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_43, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_44, negated_conjecture, (esk2_0!=k1_waybel_2(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_45, negated_conjecture, (k1_waybel_2(esk1_0,esk3_0)=k4_yellow_2(esk1_0,u1_waybel_0(esk1_0,esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_24])]), c_0_29])).
fof(c_0_46, plain, ![X22, X23]:(v2_struct_0(X22)|~l1_orders_2(X22)|(~v1_relat_1(X23)|k4_yellow_2(X22,X23)=k1_yellow_0(X22,k2_relat_1(X23)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_yellow_2])])])])).
fof(c_0_47, plain, ![X24, X25, X26]:(~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(X24,X25)))|v1_relat_1(X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_48, negated_conjecture, (m1_subset_1(u1_waybel_0(esk1_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0))))|~l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_30, c_0_28])).
cnf(c_0_49, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_50, negated_conjecture, (r2_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0)),X1)|m1_subset_1(esk4_4(esk1_0,esk3_0,X1,X1),u1_struct_0(esk1_0))|~r3_waybel_9(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_28]), c_0_34]), c_0_20]), c_0_35]), c_0_18]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_41]), c_0_42])]), c_0_43])).
cnf(c_0_51, negated_conjecture, (r3_waybel_9(esk1_0,esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_52, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_53, negated_conjecture, (k4_yellow_2(esk1_0,u1_waybel_0(esk1_0,esk3_0))!=esk2_0), inference(rw,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_54, plain, (v2_struct_0(X1)|k4_yellow_2(X1,X2)=k1_yellow_0(X1,k2_relat_1(X2))|~l1_orders_2(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_55, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_56, negated_conjecture, (m1_subset_1(u1_waybel_0(esk1_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_24])])).
fof(c_0_57, plain, ![X27, X28, X29]:(~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(X27,X28)))|k2_relset_1(X27,X28,X29)=k2_relat_1(X29)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
fof(c_0_58, plain, ![X30, X31, X32, X33, X35]:((m1_subset_1(esk5_4(X30,X31,X32,X33),u1_struct_0(X30))|~r3_waybel_9(X30,X31,X32)|X32!=X33|(~m1_subset_1(X35,u1_struct_0(X30))|(~r2_lattice3(X30,k2_relset_1(u1_struct_0(X31),u1_struct_0(X30),u1_waybel_0(X30,X31)),X35)|r3_orders_2(X30,X33,X35)))|~m1_subset_1(X33,u1_struct_0(X30))|~m1_subset_1(X32,u1_struct_0(X30))|(v2_struct_0(X31)|(~v4_orders_2(X31)|(~v7_waybel_0(X31)|~l1_waybel_0(X31,X30))))|(~v2_pre_topc(X30)|(~v8_pre_topc(X30)|(~v3_orders_2(X30)|(~v4_orders_2(X30)|(~v5_orders_2(X30)|(~v1_lattice3(X30)|(~v2_lattice3(X30)|(~v1_compts_1(X30)|~l1_waybel_9(X30))))))))))&(~v5_pre_topc(k4_waybel_1(X30,esk5_4(X30,X31,X32,X33)),X30,X30)|~r3_waybel_9(X30,X31,X32)|X32!=X33|(~m1_subset_1(X35,u1_struct_0(X30))|(~r2_lattice3(X30,k2_relset_1(u1_struct_0(X31),u1_struct_0(X30),u1_waybel_0(X30,X31)),X35)|r3_orders_2(X30,X33,X35)))|~m1_subset_1(X33,u1_struct_0(X30))|~m1_subset_1(X32,u1_struct_0(X30))|(v2_struct_0(X31)|(~v4_orders_2(X31)|(~v7_waybel_0(X31)|~l1_waybel_0(X31,X30))))|(~v2_pre_topc(X30)|(~v8_pre_topc(X30)|(~v3_orders_2(X30)|(~v4_orders_2(X30)|(~v5_orders_2(X30)|(~v1_lattice3(X30)|(~v2_lattice3(X30)|(~v1_compts_1(X30)|~l1_waybel_9(X30))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l49_waybel_9])])])])])])).
cnf(c_0_59, plain, (r2_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X4)|v2_struct_0(X2)|~v5_pre_topc(k4_waybel_1(X1,esk4_4(X1,X2,X3,X4)),X1,X1)|~v10_waybel_0(X2,X1)|~r3_waybel_9(X1,X2,X3)|X3!=X4|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v1_compts_1(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_60, negated_conjecture, (v5_pre_topc(k4_waybel_1(esk1_0,X1),esk1_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_61, negated_conjecture, (r2_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0)),esk2_0)|m1_subset_1(esk4_4(esk1_0,esk3_0,esk2_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52])])).
cnf(c_0_62, negated_conjecture, (k1_yellow_0(esk1_0,k2_relat_1(u1_waybel_0(esk1_0,esk3_0)))!=esk2_0|~v1_relat_1(u1_waybel_0(esk1_0,esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_24])]), c_0_29])).
cnf(c_0_63, negated_conjecture, (v1_relat_1(u1_waybel_0(esk1_0,esk3_0))), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_64, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_65, plain, (m1_subset_1(esk5_4(X1,X2,X3,X4),u1_struct_0(X1))|r3_orders_2(X1,X4,X5)|v2_struct_0(X2)|~r3_waybel_9(X1,X2,X3)|X3!=X4|~m1_subset_1(X5,u1_struct_0(X1))|~r2_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X5)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v1_compts_1(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
fof(c_0_66, plain, ![X36, X37, X38, X39, X40]:(((r2_lattice3(X36,X38,X37)|(X37!=k1_yellow_0(X36,X38)|~r1_yellow_0(X36,X38))|~m1_subset_1(X37,u1_struct_0(X36))|(~v5_orders_2(X36)|~l1_orders_2(X36)))&(~m1_subset_1(X39,u1_struct_0(X36))|(~r2_lattice3(X36,X38,X39)|r1_orders_2(X36,X37,X39))|(X37!=k1_yellow_0(X36,X38)|~r1_yellow_0(X36,X38))|~m1_subset_1(X37,u1_struct_0(X36))|(~v5_orders_2(X36)|~l1_orders_2(X36))))&(((X37=k1_yellow_0(X36,X40)|(m1_subset_1(esk6_3(X36,X37,X40),u1_struct_0(X36))|~r2_lattice3(X36,X40,X37))|~m1_subset_1(X37,u1_struct_0(X36))|(~v5_orders_2(X36)|~l1_orders_2(X36)))&(r1_yellow_0(X36,X40)|(m1_subset_1(esk6_3(X36,X37,X40),u1_struct_0(X36))|~r2_lattice3(X36,X40,X37))|~m1_subset_1(X37,u1_struct_0(X36))|(~v5_orders_2(X36)|~l1_orders_2(X36))))&(((X37=k1_yellow_0(X36,X40)|(r2_lattice3(X36,X40,esk6_3(X36,X37,X40))|~r2_lattice3(X36,X40,X37))|~m1_subset_1(X37,u1_struct_0(X36))|(~v5_orders_2(X36)|~l1_orders_2(X36)))&(r1_yellow_0(X36,X40)|(r2_lattice3(X36,X40,esk6_3(X36,X37,X40))|~r2_lattice3(X36,X40,X37))|~m1_subset_1(X37,u1_struct_0(X36))|(~v5_orders_2(X36)|~l1_orders_2(X36))))&((X37=k1_yellow_0(X36,X40)|(~r1_orders_2(X36,X37,esk6_3(X36,X37,X40))|~r2_lattice3(X36,X40,X37))|~m1_subset_1(X37,u1_struct_0(X36))|(~v5_orders_2(X36)|~l1_orders_2(X36)))&(r1_yellow_0(X36,X40)|(~r1_orders_2(X36,X37,esk6_3(X36,X37,X40))|~r2_lattice3(X36,X40,X37))|~m1_subset_1(X37,u1_struct_0(X36))|(~v5_orders_2(X36)|~l1_orders_2(X36))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_yellow_0])])])])])])).
cnf(c_0_67, plain, (r2_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X3)|v2_struct_0(X2)|~r3_waybel_9(X1,X2,X3)|~v10_waybel_0(X2,X1)|~v5_pre_topc(k4_waybel_1(X1,esk4_4(X1,X2,X3,X3)),X1,X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_59])).
cnf(c_0_68, negated_conjecture, (r2_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0)),esk2_0)|v5_pre_topc(k4_waybel_1(esk1_0,esk4_4(esk1_0,esk3_0,esk2_0,esk2_0)),esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_69, negated_conjecture, (k1_yellow_0(esk1_0,k2_relat_1(u1_waybel_0(esk1_0,esk3_0)))!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_63])])).
cnf(c_0_70, negated_conjecture, (k2_relat_1(u1_waybel_0(esk1_0,esk3_0))=k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0))), inference(spm,[status(thm)],[c_0_64, c_0_56])).
cnf(c_0_71, plain, (r3_orders_2(X1,X2,X3)|v2_struct_0(X4)|m1_subset_1(esk5_4(X1,X4,X2,X2),u1_struct_0(X1))|~r2_lattice3(X1,k2_relset_1(u1_struct_0(X4),u1_struct_0(X1),u1_waybel_0(X1,X4)),X3)|~r3_waybel_9(X1,X4,X2)|~l1_waybel_0(X4,X1)|~v7_waybel_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X4)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_65])).
cnf(c_0_72, plain, (X1=k1_yellow_0(X2,X3)|r2_lattice3(X2,X3,esk6_3(X2,X1,X3))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_73, negated_conjecture, (r2_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0)),esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_51]), c_0_33]), c_0_28]), c_0_34]), c_0_52]), c_0_20]), c_0_35]), c_0_18]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_41]), c_0_42])]), c_0_43])).
cnf(c_0_74, negated_conjecture, (esk2_0!=k1_yellow_0(esk1_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0)))), inference(rw,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_75, plain, (X1=k1_yellow_0(X2,X3)|m1_subset_1(esk6_3(X2,X1,X3),u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_76, negated_conjecture, (r3_orders_2(esk1_0,esk2_0,X1)|m1_subset_1(esk5_4(esk1_0,esk3_0,esk2_0,esk2_0),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0)),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_51]), c_0_28]), c_0_34]), c_0_52]), c_0_20]), c_0_35]), c_0_18]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_41]), c_0_42])]), c_0_43])).
cnf(c_0_77, negated_conjecture, (r2_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0)),esk6_3(esk1_0,esk2_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_24]), c_0_52]), c_0_37])]), c_0_74])).
cnf(c_0_78, negated_conjecture, (m1_subset_1(esk6_3(esk1_0,esk2_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0))),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_73]), c_0_24]), c_0_52]), c_0_37])]), c_0_74])).
cnf(c_0_79, plain, (r3_orders_2(X1,X4,X5)|v2_struct_0(X2)|~v5_pre_topc(k4_waybel_1(X1,esk5_4(X1,X2,X3,X4)),X1,X1)|~r3_waybel_9(X1,X2,X3)|X3!=X4|~m1_subset_1(X5,u1_struct_0(X1))|~r2_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X5)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v1_compts_1(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_80, negated_conjecture, (r3_orders_2(esk1_0,esk2_0,esk6_3(esk1_0,esk2_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0))))|m1_subset_1(esk5_4(esk1_0,esk3_0,esk2_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78])])).
cnf(c_0_81, plain, (r3_orders_2(X1,X2,X3)|v2_struct_0(X4)|~r2_lattice3(X1,k2_relset_1(u1_struct_0(X4),u1_struct_0(X1),u1_waybel_0(X1,X4)),X3)|~r3_waybel_9(X1,X4,X2)|~v5_pre_topc(k4_waybel_1(X1,esk5_4(X1,X4,X2,X2)),X1,X1)|~l1_waybel_0(X4,X1)|~v7_waybel_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X4)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_79])).
cnf(c_0_82, negated_conjecture, (r3_orders_2(esk1_0,esk2_0,esk6_3(esk1_0,esk2_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0))))|v5_pre_topc(k4_waybel_1(esk1_0,esk5_4(esk1_0,esk3_0,esk2_0,esk2_0)),esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_60, c_0_80])).
fof(c_0_83, plain, ![X42, X43, X44]:((~r3_orders_2(X42,X43,X44)|r1_orders_2(X42,X43,X44)|(v2_struct_0(X42)|(~v3_orders_2(X42)|~l1_orders_2(X42))|(~m1_subset_1(X43,u1_struct_0(X42))|~m1_subset_1(X44,u1_struct_0(X42)))))&(~r1_orders_2(X42,X43,X44)|r3_orders_2(X42,X43,X44)|(v2_struct_0(X42)|(~v3_orders_2(X42)|~l1_orders_2(X42))|(~m1_subset_1(X43,u1_struct_0(X42))|~m1_subset_1(X44,u1_struct_0(X42)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_84, negated_conjecture, (r3_orders_2(esk1_0,esk2_0,esk6_3(esk1_0,esk2_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0))))|r3_orders_2(esk1_0,esk2_0,X1)|~r2_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0)),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_51]), c_0_28]), c_0_34]), c_0_52]), c_0_20]), c_0_35]), c_0_18]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_41]), c_0_42])]), c_0_43])).
cnf(c_0_85, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_86, negated_conjecture, (r3_orders_2(esk1_0,esk2_0,esk6_3(esk1_0,esk2_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_77]), c_0_78])])).
cnf(c_0_87, plain, (X1=k1_yellow_0(X2,X3)|~r1_orders_2(X2,X1,esk6_3(X2,X1,X3))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_88, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk6_3(esk1_0,esk2_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_24]), c_0_78]), c_0_52]), c_0_40])]), c_0_29])).
cnf(c_0_89, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_73]), c_0_24]), c_0_52]), c_0_37])]), c_0_74]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 90
# Proof object clause steps            : 63
# Proof object formula steps           : 27
# Proof object conjectures             : 46
# Proof object clause conjectures      : 43
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 34
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 108
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 43
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 43
# Processed clauses                    : 130
# ...of these trivial                  : 2
# ...subsumed                          : 0
# ...remaining for further processing  : 128
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 15
# Generated clauses                    : 50
# ...of the previous two non-trivial   : 50
# Contextual simplify-reflections      : 0
# Paramodulations                      : 44
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 64
#    Positive orientable unit clauses  : 31
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 29
# Current number of unprocessed clauses: 4
# ...number of literals in the above   : 8
# Current number of archived formulas  : 0
# Current number of archived clauses   : 58
# Clause-clause subsumption calls (NU) : 4108
# Rec. Clause-clause subsumption calls : 151
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 29
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 18
# BW rewrite match successes           : 10
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6400

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.000 s
# Total time               : 0.013 s
# Maximum resident set size: 28736 pages
