# No SInE strategy applied
# Auto-Mode selected heuristic G_E___200_B02_F1_SE_CS_SP_PI_S2k
# and selection function SelectCQArNTEqFirst.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t76_funct_2, conjecture, ![X1, X2]:((v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>((r2_relset_1(X1,X1,k1_partfun1(X1,X1,X1,X1,X3,X2),X2)&v2_funct_1(X2))=>r2_relset_1(X1,X1,X3,k6_partfun1(X1))))), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', t76_funct_2)).
fof(t67_funct_2, axiom, ![X1, X2]:((v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>k1_relset_1(X1,X1,X2)=X1), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', t67_funct_2)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', redefinition_k1_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', cc1_relset_1)).
fof(t3_funct_2, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_funct_1(X1)&(v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))))), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', t3_funct_2)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', redefinition_k2_relset_1)).
fof(t31_funct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>((v2_funct_1(X3)&k2_relset_1(X1,X2,X3)=X2)=>(v1_funct_1(k2_funct_1(X3))&(v1_funct_2(k2_funct_1(X3),X2,X1)&m1_subset_1(k2_funct_1(X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1))))))), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', t31_funct_2)).
fof(redefinition_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:(((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&(v1_funct_1(X6)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))))=>k1_partfun1(X1,X2,X3,X4,X5,X6)=k5_relat_1(X5,X6)), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', redefinition_k1_partfun1)).
fof(dt_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:(((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&(v1_funct_1(X6)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))))=>(v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))&m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4))))), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', dt_k1_partfun1)).
fof(redefinition_k4_relset_1, axiom, ![X1, X2, X3, X4, X5, X6]:((m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>k4_relset_1(X1,X2,X3,X4,X5,X6)=k5_relat_1(X5,X6)), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', redefinition_k4_relset_1)).
fof(t61_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)=>(k5_relat_1(X1,k2_funct_1(X1))=k6_relat_1(k1_relat_1(X1))&k5_relat_1(k2_funct_1(X1),X1)=k6_relat_1(k2_relat_1(X1))))), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', t61_funct_1)).
fof(redefinition_k6_partfun1, axiom, ![X1]:k6_partfun1(X1)=k6_relat_1(X1), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', redefinition_k6_partfun1)).
fof(dt_k4_relset_1, axiom, ![X1, X2, X3, X4, X5, X6]:((m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>m1_subset_1(k4_relset_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', dt_k4_relset_1)).
fof(dt_k6_partfun1, axiom, ![X1]:(v1_partfun1(k6_partfun1(X1),X1)&m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', dt_k6_partfun1)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', t4_subset_1)).
fof(redefinition_r2_relset_1, axiom, ![X1, X2, X3, X4]:((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>(r2_relset_1(X1,X2,X3,X4)<=>X3=X4)), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', redefinition_r2_relset_1)).
fof(t23_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(r2_relset_1(X1,X2,k4_relset_1(X1,X1,X1,X2,k6_partfun1(X1),X3),X3)&r2_relset_1(X1,X2,k4_relset_1(X1,X2,X2,X2,X3,k6_partfun1(X2)),X3))), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', t23_funct_2)).
fof(t71_relat_1, axiom, ![X1]:(k1_relat_1(k6_relat_1(X1))=X1&k2_relat_1(k6_relat_1(X1))=X1), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', t71_relat_1)).
fof(fc4_funct_1, axiom, ![X1]:(v1_relat_1(k6_relat_1(X1))&v2_funct_1(k6_relat_1(X1))), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', fc4_funct_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', t113_zfmisc_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', t3_subset)).
fof(t27_funct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&~((v2_funct_1(X3)<=>![X4, X5]:((v1_funct_1(X5)&(v1_funct_2(X5,X4,X1)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X1)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,X4,X1)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X1)))))=>(r2_relset_1(X4,X2,k1_partfun1(X4,X1,X1,X2,X5,X3),k1_partfun1(X4,X1,X1,X2,X6,X3))=>r2_relset_1(X4,X1,X5,X6)))))))))), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', t27_funct_2)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/607a15b60447440f8694144bce6c8ebd.E_input', d10_xboole_0)).
fof(c_0_23, negated_conjecture, ~(![X1, X2]:((v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))=>((r2_relset_1(X1,X1,k1_partfun1(X1,X1,X1,X1,X3,X2),X2)&v2_funct_1(X2))=>r2_relset_1(X1,X1,X3,k6_partfun1(X1)))))), inference(assume_negation,[status(cth)],[t76_funct_2])).
fof(c_0_24, plain, ![X47, X48]:(~v1_funct_1(X48)|(~v1_funct_2(X48,X47,X47)|~m1_subset_1(X48,k1_zfmisc_1(k2_zfmisc_1(X47,X47))))|k1_relset_1(X47,X47,X48)=X47), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t67_funct_2])])).
fof(c_0_25, negated_conjecture, ((v1_funct_1(esk2_0)&(v1_funct_2(esk2_0,esk1_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,esk1_0,esk1_0)&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))))&((r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0),esk2_0)&v2_funct_1(esk2_0))&~r2_relset_1(esk1_0,esk1_0,esk3_0,k6_partfun1(esk1_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])).
fof(c_0_26, plain, ![X38, X39, X40]:(~m1_subset_1(X40,k1_zfmisc_1(k2_zfmisc_1(X38,X39)))|k1_relset_1(X38,X39,X40)=k1_relat_1(X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_27, plain, (k1_relset_1(X2,X2,X1)=X2|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_28, negated_conjecture, (v1_funct_2(esk2_0,esk1_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_31, plain, ![X30, X31, X32]:(~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(X30,X31)))|v1_relat_1(X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_32, plain, ![X57]:((v1_funct_1(X57)|(~v1_relat_1(X57)|~v1_funct_1(X57)))&((v1_funct_2(X57,k1_relat_1(X57),k2_relat_1(X57))|(~v1_relat_1(X57)|~v1_funct_1(X57)))&(m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X57),k2_relat_1(X57))))|(~v1_relat_1(X57)|~v1_funct_1(X57))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_funct_2])])])).
cnf(c_0_33, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, negated_conjecture, (k1_relset_1(esk1_0,esk1_0,esk2_0)=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30])])).
cnf(c_0_35, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_36, plain, ![X15, X16, X17]:(~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(X15,X16)))|k2_relset_1(X15,X16,X17)=k2_relat_1(X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
cnf(c_0_37, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, negated_conjecture, (k1_relat_1(esk2_0)=esk1_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_29]), c_0_34])).
cnf(c_0_39, negated_conjecture, (v1_relat_1(esk2_0)), inference(spm,[status(thm)],[c_0_35, c_0_29])).
fof(c_0_40, plain, ![X10, X11, X12]:((v1_funct_1(k2_funct_1(X12))|(~v2_funct_1(X12)|k2_relset_1(X10,X11,X12)!=X11)|(~v1_funct_1(X12)|(~v1_funct_2(X12,X10,X11)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(X10,X11))))))&((v1_funct_2(k2_funct_1(X12),X11,X10)|(~v2_funct_1(X12)|k2_relset_1(X10,X11,X12)!=X11)|(~v1_funct_1(X12)|(~v1_funct_2(X12,X10,X11)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(X10,X11))))))&(m1_subset_1(k2_funct_1(X12),k1_zfmisc_1(k2_zfmisc_1(X11,X10)))|(~v2_funct_1(X12)|k2_relset_1(X10,X11,X12)!=X11)|(~v1_funct_1(X12)|(~v1_funct_2(X12,X10,X11)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(X10,X11)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_funct_2])])])).
cnf(c_0_41, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_42, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k2_relat_1(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39]), c_0_30])])).
cnf(c_0_43, plain, (v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_44, plain, ![X24, X25, X26, X27, X28, X29]:(~v1_funct_1(X28)|~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(X24,X25)))|(~v1_funct_1(X29)|~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(X26,X27))))|k1_partfun1(X24,X25,X26,X27,X28,X29)=k5_relat_1(X28,X29)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_partfun1])])).
cnf(c_0_45, plain, (m1_subset_1(k2_funct_1(X1),k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v2_funct_1(X1)|k2_relset_1(X3,X2,X1)!=X2|~v1_funct_1(X1)|~v1_funct_2(X1,X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_46, negated_conjecture, (k2_relset_1(esk1_0,k2_relat_1(esk2_0),esk2_0)=k2_relat_1(esk2_0)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_47, negated_conjecture, (v2_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_48, negated_conjecture, (v1_funct_2(esk2_0,esk1_0,k2_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_38]), c_0_39]), c_0_30])])).
cnf(c_0_49, plain, (v1_funct_1(k2_funct_1(X1))|~v2_funct_1(X1)|k2_relset_1(X2,X3,X1)!=X3|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_50, plain, ![X41, X42, X43, X44, X45, X46]:((v1_funct_1(k1_partfun1(X41,X42,X43,X44,X45,X46))|(~v1_funct_1(X45)|~m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(X41,X42)))|(~v1_funct_1(X46)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(X43,X44))))))&(m1_subset_1(k1_partfun1(X41,X42,X43,X44,X45,X46),k1_zfmisc_1(k2_zfmisc_1(X41,X44)))|(~v1_funct_1(X45)|~m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(X41,X42)))|(~v1_funct_1(X46)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(X43,X44))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_partfun1])])])).
cnf(c_0_51, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k5_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_52, negated_conjecture, (m1_subset_1(k2_funct_1(esk2_0),k1_zfmisc_1(k2_zfmisc_1(k2_relat_1(esk2_0),esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47]), c_0_42]), c_0_48]), c_0_30])])).
cnf(c_0_53, negated_conjecture, (v1_funct_1(k2_funct_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_46]), c_0_47]), c_0_42]), c_0_48]), c_0_30])])).
fof(c_0_54, plain, ![X51, X52, X53, X54, X55, X56]:(~m1_subset_1(X55,k1_zfmisc_1(k2_zfmisc_1(X51,X52)))|~m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(X53,X54)))|k4_relset_1(X51,X52,X53,X54,X55,X56)=k5_relat_1(X55,X56)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_relset_1])])).
cnf(c_0_55, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_56, negated_conjecture, (k1_partfun1(X1,X2,k2_relat_1(esk2_0),esk1_0,X3,k2_funct_1(esk2_0))=k5_relat_1(X3,k2_funct_1(esk2_0))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53])])).
fof(c_0_57, plain, ![X58]:((k5_relat_1(X58,k2_funct_1(X58))=k6_relat_1(k1_relat_1(X58))|~v2_funct_1(X58)|(~v1_relat_1(X58)|~v1_funct_1(X58)))&(k5_relat_1(k2_funct_1(X58),X58)=k6_relat_1(k2_relat_1(X58))|~v2_funct_1(X58)|(~v1_relat_1(X58)|~v1_funct_1(X58)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t61_funct_1])])])).
fof(c_0_58, plain, ![X76]:k6_partfun1(X76)=k6_relat_1(X76), inference(variable_rename,[status(thm)],[redefinition_k6_partfun1])).
fof(c_0_59, plain, ![X18, X19, X20, X21, X22, X23]:(~m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(X18,X19)))|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(X20,X21)))|m1_subset_1(k4_relset_1(X18,X19,X20,X21,X22,X23),k1_zfmisc_1(k2_zfmisc_1(X18,X21)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_relset_1])])).
cnf(c_0_60, plain, (k4_relset_1(X2,X3,X5,X6,X1,X4)=k5_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_61, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_50])).
fof(c_0_62, plain, ![X50]:(v1_partfun1(k6_partfun1(X50),X50)&m1_subset_1(k6_partfun1(X50),k1_zfmisc_1(k2_zfmisc_1(X50,X50)))), inference(variable_rename,[status(thm)],[dt_k6_partfun1])).
cnf(c_0_63, negated_conjecture, (v1_funct_1(k1_partfun1(X1,X2,k2_relat_1(esk2_0),esk1_0,X3,k2_funct_1(esk2_0)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_52]), c_0_53])])).
cnf(c_0_64, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,k2_relat_1(esk2_0),esk1_0,esk2_0,k2_funct_1(esk2_0))=k5_relat_1(esk2_0,k2_funct_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_29]), c_0_30])])).
cnf(c_0_65, plain, (k5_relat_1(X1,k2_funct_1(X1))=k6_relat_1(k1_relat_1(X1))|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_66, plain, (k6_partfun1(X1)=k6_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
fof(c_0_67, plain, ![X69]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X69)), inference(variable_rename,[status(thm)],[t4_subset_1])).
fof(c_0_68, plain, ![X70, X71, X72, X73]:((~r2_relset_1(X70,X71,X72,X73)|X72=X73|(~m1_subset_1(X72,k1_zfmisc_1(k2_zfmisc_1(X70,X71)))|~m1_subset_1(X73,k1_zfmisc_1(k2_zfmisc_1(X70,X71)))))&(X72!=X73|r2_relset_1(X70,X71,X72,X73)|(~m1_subset_1(X72,k1_zfmisc_1(k2_zfmisc_1(X70,X71)))|~m1_subset_1(X73,k1_zfmisc_1(k2_zfmisc_1(X70,X71)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r2_relset_1])])])).
cnf(c_0_69, plain, (m1_subset_1(k4_relset_1(X2,X3,X5,X6,X1,X4),k1_zfmisc_1(k2_zfmisc_1(X2,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_70, negated_conjecture, (k4_relset_1(X1,X2,esk1_0,esk1_0,X3,esk2_0)=k5_relat_1(X3,esk2_0)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_60, c_0_29])).
cnf(c_0_71, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_72, plain, ![X35, X36, X37]:((r2_relset_1(X35,X36,k4_relset_1(X35,X35,X35,X36,k6_partfun1(X35),X37),X37)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(X35,X36))))&(r2_relset_1(X35,X36,k4_relset_1(X35,X36,X36,X36,X37,k6_partfun1(X36)),X37)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(X35,X36))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_funct_2])])])).
cnf(c_0_73, negated_conjecture, (m1_subset_1(k1_partfun1(X1,X2,esk1_0,esk1_0,X3,esk2_0),k1_zfmisc_1(k2_zfmisc_1(X1,esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_29]), c_0_30])])).
cnf(c_0_74, plain, (m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_75, negated_conjecture, (v1_funct_1(k5_relat_1(esk2_0,k2_funct_1(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_29]), c_0_64]), c_0_30])])).
cnf(c_0_76, plain, (k5_relat_1(X1,k2_funct_1(X1))=k6_partfun1(k1_relat_1(X1))|~v1_funct_1(X1)|~v2_funct_1(X1)|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_77, negated_conjecture, (k1_partfun1(X1,X2,esk1_0,esk1_0,X3,esk2_0)=k5_relat_1(X3,esk2_0)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_29]), c_0_30])])).
cnf(c_0_78, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_79, plain, (X3=X4|~r2_relset_1(X1,X2,X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_80, negated_conjecture, (r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_81, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_82, negated_conjecture, (m1_subset_1(k4_relset_1(X1,X2,esk1_0,esk1_0,X3,esk2_0),k1_zfmisc_1(k2_zfmisc_1(X1,esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_69, c_0_29])).
cnf(c_0_83, negated_conjecture, (k4_relset_1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)=k5_relat_1(esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_84, plain, (r2_relset_1(X1,X2,k4_relset_1(X1,X1,X1,X2,k6_partfun1(X1),X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_85, negated_conjecture, (m1_subset_1(k1_partfun1(X1,X1,esk1_0,esk1_0,k6_partfun1(X1),esk2_0),k1_zfmisc_1(k2_zfmisc_1(X1,esk1_0)))|~v1_funct_1(k6_partfun1(X1))), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_86, negated_conjecture, (v1_funct_1(k6_partfun1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_38]), c_0_39]), c_0_47]), c_0_30])])).
cnf(c_0_87, negated_conjecture, (k1_partfun1(X1,X1,esk1_0,esk1_0,k6_partfun1(X1),esk2_0)=k5_relat_1(k6_partfun1(X1),esk2_0)|~v1_funct_1(k6_partfun1(X1))), inference(spm,[status(thm)],[c_0_77, c_0_74])).
fof(c_0_88, plain, ![X68]:(k1_relat_1(k6_relat_1(X68))=X68&k2_relat_1(k6_relat_1(X68))=X68), inference(variable_rename,[status(thm)],[t71_relat_1])).
fof(c_0_89, plain, ![X49]:(v1_relat_1(k6_relat_1(X49))&v2_funct_1(k6_relat_1(X49))), inference(variable_rename,[status(thm)],[fc4_funct_1])).
cnf(c_0_90, plain, (k4_relset_1(X1,X2,X3,X4,X5,k1_xboole_0)=k5_relat_1(X5,k1_xboole_0)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_60, c_0_78])).
fof(c_0_91, plain, ![X33, X34]:((k2_zfmisc_1(X33,X34)!=k1_xboole_0|(X33=k1_xboole_0|X34=k1_xboole_0))&((X33!=k1_xboole_0|k2_zfmisc_1(X33,X34)=k1_xboole_0)&(X34!=k1_xboole_0|k2_zfmisc_1(X33,X34)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_92, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X3,X4,k6_partfun1(X3)),k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(k6_partfun1(X3))|~v1_funct_1(X4)), inference(spm,[status(thm)],[c_0_61, c_0_74])).
cnf(c_0_93, plain, (k1_partfun1(X1,X2,X3,X3,X4,k6_partfun1(X3))=k5_relat_1(X4,k6_partfun1(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(k6_partfun1(X3))|~v1_funct_1(X4)), inference(spm,[status(thm)],[c_0_51, c_0_74])).
cnf(c_0_94, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)=esk2_0|~m1_subset_1(k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_29])])).
cnf(c_0_95, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)=k5_relat_1(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_71]), c_0_81])])).
cnf(c_0_96, negated_conjecture, (m1_subset_1(k5_relat_1(esk3_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_71]), c_0_83])).
cnf(c_0_97, negated_conjecture, (r2_relset_1(esk1_0,esk1_0,k4_relset_1(esk1_0,esk1_0,esk1_0,esk1_0,k6_partfun1(esk1_0),esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_84, c_0_29])).
cnf(c_0_98, negated_conjecture, (m1_subset_1(k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,k6_partfun1(esk1_0),esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_99, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,k6_partfun1(esk1_0),esk2_0)=k5_relat_1(k6_partfun1(esk1_0),esk2_0)), inference(spm,[status(thm)],[c_0_87, c_0_86])).
cnf(c_0_100, plain, (k2_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_101, plain, (k1_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_102, plain, (v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_89])).
fof(c_0_103, plain, ![X13, X14]:((~m1_subset_1(X13,k1_zfmisc_1(X14))|r1_tarski(X13,X14))&(~r1_tarski(X13,X14)|m1_subset_1(X13,k1_zfmisc_1(X14)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_104, plain, (m1_subset_1(k4_relset_1(X1,X2,X3,X4,X5,k1_xboole_0),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_69, c_0_78])).
cnf(c_0_105, negated_conjecture, (k4_relset_1(esk1_0,esk1_0,X1,X2,esk2_0,k1_xboole_0)=k5_relat_1(esk2_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_90, c_0_29])).
cnf(c_0_106, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_107, plain, (r2_relset_1(X1,X2,k4_relset_1(X1,X2,X2,X2,X3,k6_partfun1(X2)),X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_108, negated_conjecture, (m1_subset_1(k1_partfun1(esk1_0,esk1_0,X1,X1,esk2_0,k6_partfun1(X1)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,X1)))|~v1_funct_1(k6_partfun1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_29]), c_0_30])])).
cnf(c_0_109, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,X1,X1,esk2_0,k6_partfun1(X1))=k5_relat_1(esk2_0,k6_partfun1(X1))|~v1_funct_1(k6_partfun1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_29]), c_0_30])])).
fof(c_0_110, plain, ![X59, X60, X61, X62, X63, X64]:((~v2_funct_1(X61)|(~v1_funct_1(X63)|(~v1_funct_2(X63,X62,X59)|~m1_subset_1(X63,k1_zfmisc_1(k2_zfmisc_1(X62,X59))))|(~v1_funct_1(X64)|(~v1_funct_2(X64,X62,X59)|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(X62,X59))))|(~r2_relset_1(X62,X60,k1_partfun1(X62,X59,X59,X60,X63,X61),k1_partfun1(X62,X59,X59,X60,X64,X61))|r2_relset_1(X62,X59,X63,X64))))|X60=k1_xboole_0|X59=k1_xboole_0|(~v1_funct_1(X61)|(~v1_funct_2(X61,X59,X60)|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(X59,X60))))))&(((v1_funct_1(esk5_3(X59,X60,X61))|v2_funct_1(X61)|X60=k1_xboole_0|X59=k1_xboole_0|(~v1_funct_1(X61)|(~v1_funct_2(X61,X59,X60)|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(X59,X60))))))&((v1_funct_2(esk5_3(X59,X60,X61),esk4_3(X59,X60,X61),X59)|v2_funct_1(X61)|X60=k1_xboole_0|X59=k1_xboole_0|(~v1_funct_1(X61)|(~v1_funct_2(X61,X59,X60)|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(X59,X60))))))&(m1_subset_1(esk5_3(X59,X60,X61),k1_zfmisc_1(k2_zfmisc_1(esk4_3(X59,X60,X61),X59)))|v2_funct_1(X61)|X60=k1_xboole_0|X59=k1_xboole_0|(~v1_funct_1(X61)|(~v1_funct_2(X61,X59,X60)|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(X59,X60))))))))&(((v1_funct_1(esk6_3(X59,X60,X61))|v2_funct_1(X61)|X60=k1_xboole_0|X59=k1_xboole_0|(~v1_funct_1(X61)|(~v1_funct_2(X61,X59,X60)|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(X59,X60))))))&((v1_funct_2(esk6_3(X59,X60,X61),esk4_3(X59,X60,X61),X59)|v2_funct_1(X61)|X60=k1_xboole_0|X59=k1_xboole_0|(~v1_funct_1(X61)|(~v1_funct_2(X61,X59,X60)|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(X59,X60))))))&(m1_subset_1(esk6_3(X59,X60,X61),k1_zfmisc_1(k2_zfmisc_1(esk4_3(X59,X60,X61),X59)))|v2_funct_1(X61)|X60=k1_xboole_0|X59=k1_xboole_0|(~v1_funct_1(X61)|(~v1_funct_2(X61,X59,X60)|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(X59,X60))))))))&((r2_relset_1(esk4_3(X59,X60,X61),X60,k1_partfun1(esk4_3(X59,X60,X61),X59,X59,X60,esk5_3(X59,X60,X61),X61),k1_partfun1(esk4_3(X59,X60,X61),X59,X59,X60,esk6_3(X59,X60,X61),X61))|v2_funct_1(X61)|X60=k1_xboole_0|X59=k1_xboole_0|(~v1_funct_1(X61)|(~v1_funct_2(X61,X59,X60)|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(X59,X60))))))&(~r2_relset_1(esk4_3(X59,X60,X61),X59,esk5_3(X59,X60,X61),esk6_3(X59,X60,X61))|v2_funct_1(X61)|X60=k1_xboole_0|X59=k1_xboole_0|(~v1_funct_1(X61)|(~v1_funct_2(X61,X59,X60)|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(X59,X60)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t27_funct_2])])])])])])).
cnf(c_0_111, negated_conjecture, (k5_relat_1(esk3_0,esk2_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_94, c_0_95]), c_0_95]), c_0_96])])).
cnf(c_0_112, negated_conjecture, (k4_relset_1(esk1_0,esk1_0,esk1_0,esk1_0,k6_partfun1(esk1_0),esk2_0)=esk2_0|~m1_subset_1(k4_relset_1(esk1_0,esk1_0,esk1_0,esk1_0,k6_partfun1(esk1_0),esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_97]), c_0_29])])).
cnf(c_0_113, negated_conjecture, (k4_relset_1(X1,X1,esk1_0,esk1_0,k6_partfun1(X1),esk2_0)=k5_relat_1(k6_partfun1(X1),esk2_0)), inference(spm,[status(thm)],[c_0_70, c_0_74])).
cnf(c_0_114, negated_conjecture, (m1_subset_1(k5_relat_1(k6_partfun1(esk1_0),esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(rw,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_115, plain, (r2_relset_1(X3,X4,X1,X2)|X1!=X2|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_116, plain, (k2_relat_1(k6_partfun1(X1))=X1), inference(rw,[status(thm)],[c_0_100, c_0_66])).
cnf(c_0_117, plain, (k1_relat_1(k6_partfun1(X1))=X1), inference(rw,[status(thm)],[c_0_101, c_0_66])).
cnf(c_0_118, plain, (v1_relat_1(k6_partfun1(X1))), inference(rw,[status(thm)],[c_0_102, c_0_66])).
cnf(c_0_119, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_120, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_121, negated_conjecture, (m1_subset_1(k5_relat_1(esk2_0,k1_xboole_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,X1)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_29]), c_0_105])).
cnf(c_0_122, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_106])).
cnf(c_0_123, negated_conjecture, (k4_relset_1(esk1_0,esk1_0,X1,X2,esk3_0,k1_xboole_0)=k5_relat_1(esk3_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_90, c_0_71])).
cnf(c_0_124, negated_conjecture, (r2_relset_1(esk1_0,esk1_0,k4_relset_1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,k6_partfun1(esk1_0)),esk2_0)), inference(spm,[status(thm)],[c_0_107, c_0_29])).
cnf(c_0_125, plain, (k4_relset_1(X1,X2,X3,X3,X4,k6_partfun1(X3))=k5_relat_1(X4,k6_partfun1(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_60, c_0_74])).
cnf(c_0_126, negated_conjecture, (m1_subset_1(k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,k6_partfun1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(spm,[status(thm)],[c_0_108, c_0_86])).
cnf(c_0_127, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,k6_partfun1(esk1_0))=k5_relat_1(esk2_0,k6_partfun1(esk1_0))), inference(spm,[status(thm)],[c_0_109, c_0_86])).
cnf(c_0_128, plain, (r2_relset_1(X3,X4,X2,X5)|X6=k1_xboole_0|X4=k1_xboole_0|~v2_funct_1(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X5)|~v1_funct_2(X5,X3,X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~r2_relset_1(X3,X6,k1_partfun1(X3,X4,X4,X6,X2,X1),k1_partfun1(X3,X4,X4,X6,X5,X1))|~v1_funct_1(X1)|~v1_funct_2(X1,X4,X6)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X6)))), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_129, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)=esk2_0), inference(rw,[status(thm)],[c_0_95, c_0_111])).
cnf(c_0_130, negated_conjecture, (v1_funct_2(esk3_0,esk1_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_131, negated_conjecture, (k5_relat_1(k6_partfun1(esk1_0),esk2_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_112, c_0_113]), c_0_113]), c_0_114])])).
cnf(c_0_132, plain, (r2_relset_1(X1,X2,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(er,[status(thm)],[c_0_115])).
cnf(c_0_133, plain, (v1_funct_2(k6_partfun1(X1),X1,X1)|~v1_funct_1(k6_partfun1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_116]), c_0_117]), c_0_118])])).
fof(c_0_134, plain, ![X74, X75]:(((r1_tarski(X74,X75)|X74!=X75)&(r1_tarski(X75,X74)|X74!=X75))&(~r1_tarski(X74,X75)|~r1_tarski(X75,X74)|X74=X75)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_135, plain, (r1_tarski(k6_partfun1(X1),k2_zfmisc_1(X1,X1))), inference(spm,[status(thm)],[c_0_119, c_0_74])).
cnf(c_0_136, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_120])).
cnf(c_0_137, negated_conjecture, (m1_subset_1(k5_relat_1(esk2_0,k1_xboole_0),k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_121, c_0_122])).
cnf(c_0_138, negated_conjecture, (m1_subset_1(k1_partfun1(esk1_0,esk1_0,X1,X1,esk3_0,k6_partfun1(X1)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,X1)))|~v1_funct_1(k6_partfun1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_71]), c_0_81])])).
cnf(c_0_139, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,X1,X1,esk3_0,k6_partfun1(X1))=k5_relat_1(esk3_0,k6_partfun1(X1))|~v1_funct_1(k6_partfun1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_71]), c_0_81])])).
cnf(c_0_140, negated_conjecture, (m1_subset_1(k5_relat_1(esk3_0,k1_xboole_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,X1)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_71]), c_0_123])).
cnf(c_0_141, negated_conjecture, (k4_relset_1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,k6_partfun1(esk1_0))=esk2_0|~m1_subset_1(k4_relset_1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,k6_partfun1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_124]), c_0_29])])).
cnf(c_0_142, negated_conjecture, (k4_relset_1(esk1_0,esk1_0,X1,X1,esk2_0,k6_partfun1(X1))=k5_relat_1(esk2_0,k6_partfun1(X1))), inference(spm,[status(thm)],[c_0_125, c_0_29])).
cnf(c_0_143, negated_conjecture, (m1_subset_1(k5_relat_1(esk2_0,k6_partfun1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(rw,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_144, negated_conjecture, (esk1_0=k1_xboole_0|r2_relset_1(esk1_0,esk1_0,esk3_0,X1)|~r2_relset_1(esk1_0,esk1_0,esk2_0,k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))|~v1_funct_2(X1,esk1_0,esk1_0)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129]), c_0_47]), c_0_71]), c_0_29]), c_0_130]), c_0_28]), c_0_81]), c_0_30])])).
cnf(c_0_145, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,k6_partfun1(esk1_0),esk2_0)=esk2_0), inference(rw,[status(thm)],[c_0_99, c_0_131])).
cnf(c_0_146, negated_conjecture, (r2_relset_1(esk1_0,esk1_0,esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_132, c_0_29])).
cnf(c_0_147, negated_conjecture, (v1_funct_2(k6_partfun1(esk1_0),esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_133, c_0_86])).
cnf(c_0_148, negated_conjecture, (~r2_relset_1(esk1_0,esk1_0,esk3_0,k6_partfun1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_149, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_150, plain, (r1_tarski(k6_partfun1(k1_xboole_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_151, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_119, c_0_78])).
cnf(c_0_152, negated_conjecture, (r1_tarski(k5_relat_1(esk2_0,k1_xboole_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_119, c_0_137])).
cnf(c_0_153, negated_conjecture, (r2_relset_1(esk1_0,esk1_0,k4_relset_1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,k6_partfun1(esk1_0)),esk3_0)), inference(spm,[status(thm)],[c_0_107, c_0_71])).
cnf(c_0_154, negated_conjecture, (m1_subset_1(k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,k6_partfun1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(spm,[status(thm)],[c_0_138, c_0_86])).
cnf(c_0_155, negated_conjecture, (k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,k6_partfun1(esk1_0))=k5_relat_1(esk3_0,k6_partfun1(esk1_0))), inference(spm,[status(thm)],[c_0_139, c_0_86])).
cnf(c_0_156, negated_conjecture, (m1_subset_1(k5_relat_1(esk3_0,k1_xboole_0),k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_140, c_0_122])).
cnf(c_0_157, negated_conjecture, (k5_relat_1(esk2_0,k6_partfun1(esk1_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_141, c_0_142]), c_0_142]), c_0_143])])).
cnf(c_0_158, negated_conjecture, (esk1_0=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_145]), c_0_146]), c_0_74]), c_0_147]), c_0_86])]), c_0_148])).
cnf(c_0_159, plain, (k6_partfun1(k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_150]), c_0_151])])).
cnf(c_0_160, negated_conjecture, (k5_relat_1(esk2_0,k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_152]), c_0_151])])).
cnf(c_0_161, negated_conjecture, (k4_relset_1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,k6_partfun1(esk1_0))=esk3_0|~m1_subset_1(k4_relset_1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,k6_partfun1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_153]), c_0_71])])).
cnf(c_0_162, negated_conjecture, (k4_relset_1(esk1_0,esk1_0,X1,X1,esk3_0,k6_partfun1(X1))=k5_relat_1(esk3_0,k6_partfun1(X1))), inference(spm,[status(thm)],[c_0_125, c_0_71])).
cnf(c_0_163, negated_conjecture, (m1_subset_1(k5_relat_1(esk3_0,k6_partfun1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(rw,[status(thm)],[c_0_154, c_0_155])).
cnf(c_0_164, negated_conjecture, (r1_tarski(k5_relat_1(esk3_0,k1_xboole_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_119, c_0_156])).
cnf(c_0_165, negated_conjecture, (k1_xboole_0=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_157, c_0_158]), c_0_159]), c_0_160])).
cnf(c_0_166, negated_conjecture, (k5_relat_1(esk3_0,k6_partfun1(esk1_0))=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_161, c_0_162]), c_0_162]), c_0_163])])).
cnf(c_0_167, negated_conjecture, (k5_relat_1(esk3_0,k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_164]), c_0_151])])).
cnf(c_0_168, plain, (r2_relset_1(X1,X1,k6_partfun1(X1),k6_partfun1(X1))), inference(spm,[status(thm)],[c_0_132, c_0_74])).
cnf(c_0_169, plain, (k6_partfun1(esk2_0)=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_159, c_0_165]), c_0_165])).
cnf(c_0_170, negated_conjecture, (esk3_0=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_166, c_0_158]), c_0_159]), c_0_167]), c_0_165])).
cnf(c_0_171, plain, (r2_relset_1(esk2_0,esk2_0,esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_168, c_0_169])).
cnf(c_0_172, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_148, c_0_158]), c_0_158]), c_0_158]), c_0_159]), c_0_165]), c_0_165]), c_0_170]), c_0_165]), c_0_171])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 173
# Proof object clause steps            : 126
# Proof object formula steps           : 47
# Proof object conjectures             : 80
# Proof object clause conjectures      : 77
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 38
# Proof object initial formulas used   : 23
# Proof object generating inferences   : 68
# Proof object simplifying inferences  : 129
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 23
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 55
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 53
# Processed clauses                    : 3629
# ...of these trivial                  : 141
# ...subsumed                          : 772
# ...remaining for further processing  : 2715
# Other redundant clauses eliminated   : 17
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 17
# Backward-rewritten                   : 2540
# Generated clauses                    : 38326
# ...of the previous two non-trivial   : 38421
# Contextual simplify-reflections      : 106
# Paramodulations                      : 38309
# Factorizations                       : 0
# Equation resolutions                 : 17
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 153
#    Positive orientable unit clauses  : 57
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 96
# Current number of unprocessed clauses: 34054
# ...number of literals in the above   : 77893
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2558
# Clause-clause subsumption calls (NU) : 53856
# Rec. Clause-clause subsumption calls : 22414
# Non-unit clause-clause subsumptions  : 895
# Unit Clause-clause subsumption calls : 20910
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 122679
# BW rewrite match successes           : 148
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1399963

# -------------------------------------------------
# User time                : 0.481 s
# System time              : 0.010 s
# Total time               : 0.491 s
# Maximum resident set size: 36792 pages
