# No SInE strategy applied
# Auto-Mode selected heuristic G_E___200_B02_F1_SE_CS_SP_PI_S2k
# and selection function SelectCQArNTEqFirst.
#
# Preprocessing time       : 0.010 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t68_tex_2, conjecture, ![X1]:((v2_pre_topc(X1)&(v1_tdlat_3(X1)&l1_pre_topc(X1)))=>![X2]:((v2_pre_topc(X2)&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>v5_pre_topc(X3,X1,X2)))), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', t68_tex_2)).
fof(t17_tdlat_3, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>(v1_tdlat_3(X1)<=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>v3_pre_topc(X2,X1)))), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', t17_tdlat_3)).
fof(cc1_tdlat_3, axiom, ![X1]:(l1_pre_topc(X1)=>(v1_tdlat_3(X1)=>v2_pre_topc(X1))), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', cc1_tdlat_3)).
fof(dt_k8_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k8_relset_1(X1,X2,X3,X4),k1_zfmisc_1(X1))), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', dt_k8_relset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', t3_subset)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', d10_xboole_0)).
fof(t55_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>((k2_struct_0(X2)=k1_xboole_0=>k2_struct_0(X1)=k1_xboole_0)=>(v5_pre_topc(X3,X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(v3_pre_topc(X4,X2)=>v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X1)))))))), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', t55_tops_2)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', t113_zfmisc_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', redefinition_k1_relset_1)).
fof(t8_boole, axiom, ![X1, X2]:~((v1_xboole_0(X1)&(~(X1=X2)&v1_xboole_0(X2)))), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', t8_boole)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', d3_struct_0)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', fc6_relat_1)).
fof(t38_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(k7_relset_1(X1,X2,X3,X1)=k2_relset_1(X1,X2,X3)&k8_relset_1(X1,X2,X3,X2)=k1_relset_1(X1,X2,X3))), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', t38_relset_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', t2_xboole_1)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', dt_l1_pre_topc)).
fof(t56_tops_2, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((v2_pre_topc(X2)&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v5_pre_topc(X3,X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>r1_tarski(k2_pre_topc(X1,k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4)),k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,k2_pre_topc(X2,X4)))))))), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', t56_tops_2)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', d4_subset_1)).
fof(d18_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(v4_relat_1(X2,X1)<=>r1_tarski(k1_relat_1(X2),X1))), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', d18_relat_1)).
fof(l222_relat_1, axiom, ![X1, X2]:(v4_relat_1(k1_xboole_0,X1)&v5_relat_1(k1_xboole_0,X2)), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', l222_relat_1)).
fof(t18_tdlat_3, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>(v1_tdlat_3(X1)<=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>v4_pre_topc(X2,X1)))), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', t18_tdlat_3)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', fc1_xboole_0)).
fof(t3_xboole_1, axiom, ![X1]:(r1_tarski(X1,k1_xboole_0)=>X1=k1_xboole_0), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', t3_xboole_1)).
fof(t52_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v4_pre_topc(X2,X1)=>k2_pre_topc(X1,X2)=X2)&((v2_pre_topc(X1)&k2_pre_topc(X1,X2)=X2)=>v4_pre_topc(X2,X1))))), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', t52_pre_topc)).
fof(cc4_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/610157f7a56740368b322bd8b1aa89d4.E_input', cc4_relset_1)).
fof(c_0_25, negated_conjecture, ~(![X1]:((v2_pre_topc(X1)&(v1_tdlat_3(X1)&l1_pre_topc(X1)))=>![X2]:((v2_pre_topc(X2)&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>v5_pre_topc(X3,X1,X2))))), inference(assume_negation,[status(cth)],[t68_tex_2])).
fof(c_0_26, plain, ![X31, X32]:((~v1_tdlat_3(X31)|(~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|v3_pre_topc(X32,X31))|(~v2_pre_topc(X31)|~l1_pre_topc(X31)))&((m1_subset_1(esk7_1(X31),k1_zfmisc_1(u1_struct_0(X31)))|v1_tdlat_3(X31)|(~v2_pre_topc(X31)|~l1_pre_topc(X31)))&(~v3_pre_topc(esk7_1(X31),X31)|v1_tdlat_3(X31)|(~v2_pre_topc(X31)|~l1_pre_topc(X31))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t17_tdlat_3])])])])])).
fof(c_0_27, plain, ![X43]:(~l1_pre_topc(X43)|(~v1_tdlat_3(X43)|v2_pre_topc(X43))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_tdlat_3])])).
fof(c_0_28, plain, ![X63, X64, X65, X66]:(~m1_subset_1(X65,k1_zfmisc_1(k2_zfmisc_1(X63,X64)))|m1_subset_1(k8_relset_1(X63,X64,X65,X66),k1_zfmisc_1(X63))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_relset_1])])).
fof(c_0_29, negated_conjecture, ((v2_pre_topc(esk1_0)&(v1_tdlat_3(esk1_0)&l1_pre_topc(esk1_0)))&((v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))))&~v5_pre_topc(esk3_0,esk1_0,esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])).
fof(c_0_30, plain, ![X56, X57]:((~m1_subset_1(X56,k1_zfmisc_1(X57))|r1_tarski(X56,X57))&(~r1_tarski(X56,X57)|m1_subset_1(X56,k1_zfmisc_1(X57)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_31, plain, (v3_pre_topc(X2,X1)|~v1_tdlat_3(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, plain, (v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, plain, (m1_subset_1(k8_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_35, plain, ![X8, X9]:(((r1_tarski(X8,X9)|X8!=X9)&(r1_tarski(X9,X8)|X8!=X9))&(~r1_tarski(X8,X9)|~r1_tarski(X9,X8)|X8=X9)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_36, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_37, plain, ![X17, X18, X19, X20]:(((~v5_pre_topc(X19,X17,X18)|(~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X18)))|(~v3_pre_topc(X20,X18)|v3_pre_topc(k8_relset_1(u1_struct_0(X17),u1_struct_0(X18),X19,X20),X17)))|k2_struct_0(X18)=k1_xboole_0|(~v1_funct_1(X19)|(~v1_funct_2(X19,u1_struct_0(X17),u1_struct_0(X18))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X17),u1_struct_0(X18))))))|~l1_pre_topc(X18)|~l1_pre_topc(X17))&((m1_subset_1(esk5_3(X17,X18,X19),k1_zfmisc_1(u1_struct_0(X18)))|v5_pre_topc(X19,X17,X18)|k2_struct_0(X18)=k1_xboole_0|(~v1_funct_1(X19)|(~v1_funct_2(X19,u1_struct_0(X17),u1_struct_0(X18))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X17),u1_struct_0(X18))))))|~l1_pre_topc(X18)|~l1_pre_topc(X17))&((v3_pre_topc(esk5_3(X17,X18,X19),X18)|v5_pre_topc(X19,X17,X18)|k2_struct_0(X18)=k1_xboole_0|(~v1_funct_1(X19)|(~v1_funct_2(X19,u1_struct_0(X17),u1_struct_0(X18))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X17),u1_struct_0(X18))))))|~l1_pre_topc(X18)|~l1_pre_topc(X17))&(~v3_pre_topc(k8_relset_1(u1_struct_0(X17),u1_struct_0(X18),X19,esk5_3(X17,X18,X19)),X17)|v5_pre_topc(X19,X17,X18)|k2_struct_0(X18)=k1_xboole_0|(~v1_funct_1(X19)|(~v1_funct_2(X19,u1_struct_0(X17),u1_struct_0(X18))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X17),u1_struct_0(X18))))))|~l1_pre_topc(X18)|~l1_pre_topc(X17)))))&((~v5_pre_topc(X19,X17,X18)|(~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X18)))|(~v3_pre_topc(X20,X18)|v3_pre_topc(k8_relset_1(u1_struct_0(X17),u1_struct_0(X18),X19,X20),X17)))|k2_struct_0(X17)!=k1_xboole_0|(~v1_funct_1(X19)|(~v1_funct_2(X19,u1_struct_0(X17),u1_struct_0(X18))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X17),u1_struct_0(X18))))))|~l1_pre_topc(X18)|~l1_pre_topc(X17))&((m1_subset_1(esk5_3(X17,X18,X19),k1_zfmisc_1(u1_struct_0(X18)))|v5_pre_topc(X19,X17,X18)|k2_struct_0(X17)!=k1_xboole_0|(~v1_funct_1(X19)|(~v1_funct_2(X19,u1_struct_0(X17),u1_struct_0(X18))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X17),u1_struct_0(X18))))))|~l1_pre_topc(X18)|~l1_pre_topc(X17))&((v3_pre_topc(esk5_3(X17,X18,X19),X18)|v5_pre_topc(X19,X17,X18)|k2_struct_0(X17)!=k1_xboole_0|(~v1_funct_1(X19)|(~v1_funct_2(X19,u1_struct_0(X17),u1_struct_0(X18))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X17),u1_struct_0(X18))))))|~l1_pre_topc(X18)|~l1_pre_topc(X17))&(~v3_pre_topc(k8_relset_1(u1_struct_0(X17),u1_struct_0(X18),X19,esk5_3(X17,X18,X19)),X17)|v5_pre_topc(X19,X17,X18)|k2_struct_0(X17)!=k1_xboole_0|(~v1_funct_1(X19)|(~v1_funct_2(X19,u1_struct_0(X17),u1_struct_0(X18))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X17),u1_struct_0(X18))))))|~l1_pre_topc(X18)|~l1_pre_topc(X17)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_tops_2])])])])])).
cnf(c_0_38, plain, (v3_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v1_tdlat_3(X2)), inference(csr,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_39, negated_conjecture, (m1_subset_1(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_40, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_41, negated_conjecture, (v1_tdlat_3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_42, plain, ![X44, X45]:((k2_zfmisc_1(X44,X45)!=k1_xboole_0|(X44=k1_xboole_0|X45=k1_xboole_0))&((X44!=k1_xboole_0|k2_zfmisc_1(X44,X45)=k1_xboole_0)&(X45!=k1_xboole_0|k2_zfmisc_1(X44,X45)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_43, plain, ![X22, X23, X24]:(~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(X22,X23)))|k1_relset_1(X22,X23,X24)=k1_relat_1(X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_44, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_45, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_36, c_0_34])).
fof(c_0_46, plain, ![X61, X62]:(~v1_xboole_0(X61)|(X61=X62|~v1_xboole_0(X62))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_boole])])])).
fof(c_0_47, plain, ![X46]:(~l1_struct_0(X46)|k2_struct_0(X46)=u1_struct_0(X46)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_48, plain, (v5_pre_topc(X3,X1,X2)|k2_struct_0(X2)=k1_xboole_0|~v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk5_3(X1,X2,X3)),X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_49, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_50, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_51, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_52, negated_conjecture, (~v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_53, negated_conjecture, (v3_pre_topc(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40]), c_0_41])])).
fof(c_0_54, plain, ![X25, X26]:v1_relat_1(k2_zfmisc_1(X25,X26)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_55, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_56, plain, ![X52, X53, X54]:((k7_relset_1(X52,X53,X54,X52)=k2_relset_1(X52,X53,X54)|~m1_subset_1(X54,k1_zfmisc_1(k2_zfmisc_1(X52,X53))))&(k8_relset_1(X52,X53,X54,X53)=k1_relset_1(X52,X53,X54)|~m1_subset_1(X54,k1_zfmisc_1(k2_zfmisc_1(X52,X53))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t38_relset_1])])])).
cnf(c_0_57, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_58, negated_conjecture, (k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))=esk3_0|~r1_tarski(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)),esk3_0)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_59, plain, (X1=X2|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_60, plain, ![X35]:r1_tarski(k1_xboole_0,X35), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_61, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_62, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_34]), c_0_50]), c_0_51]), c_0_40])]), c_0_52]), c_0_53])])).
fof(c_0_63, plain, ![X58]:(~l1_pre_topc(X58)|l1_struct_0(X58)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_64, plain, ![X67, X68, X69, X70]:((~v5_pre_topc(X69,X67,X68)|(~m1_subset_1(X70,k1_zfmisc_1(u1_struct_0(X68)))|r1_tarski(k2_pre_topc(X67,k8_relset_1(u1_struct_0(X67),u1_struct_0(X68),X69,X70)),k8_relset_1(u1_struct_0(X67),u1_struct_0(X68),X69,k2_pre_topc(X68,X70))))|(~v1_funct_1(X69)|(~v1_funct_2(X69,u1_struct_0(X67),u1_struct_0(X68))|~m1_subset_1(X69,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X67),u1_struct_0(X68))))))|(~v2_pre_topc(X68)|~l1_pre_topc(X68))|(~v2_pre_topc(X67)|~l1_pre_topc(X67)))&((m1_subset_1(esk8_3(X67,X68,X69),k1_zfmisc_1(u1_struct_0(X68)))|v5_pre_topc(X69,X67,X68)|(~v1_funct_1(X69)|(~v1_funct_2(X69,u1_struct_0(X67),u1_struct_0(X68))|~m1_subset_1(X69,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X67),u1_struct_0(X68))))))|(~v2_pre_topc(X68)|~l1_pre_topc(X68))|(~v2_pre_topc(X67)|~l1_pre_topc(X67)))&(~r1_tarski(k2_pre_topc(X67,k8_relset_1(u1_struct_0(X67),u1_struct_0(X68),X69,esk8_3(X67,X68,X69))),k8_relset_1(u1_struct_0(X67),u1_struct_0(X68),X69,k2_pre_topc(X68,esk8_3(X67,X68,X69))))|v5_pre_topc(X69,X67,X68)|(~v1_funct_1(X69)|(~v1_funct_2(X69,u1_struct_0(X67),u1_struct_0(X68))|~m1_subset_1(X69,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X67),u1_struct_0(X68))))))|(~v2_pre_topc(X68)|~l1_pre_topc(X68))|(~v2_pre_topc(X67)|~l1_pre_topc(X67))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t56_tops_2])])])])])).
fof(c_0_65, plain, ![X34]:m1_subset_1(k2_subset_1(X34),k1_zfmisc_1(X34)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_66, plain, ![X36]:k2_subset_1(X36)=X36, inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_67, plain, ![X47, X48]:((~v4_relat_1(X48,X47)|r1_tarski(k1_relat_1(X48),X47)|~v1_relat_1(X48))&(~r1_tarski(k1_relat_1(X48),X47)|v4_relat_1(X48,X47)|~v1_relat_1(X48))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d18_relat_1])])])).
fof(c_0_68, plain, ![X41, X42]:(v4_relat_1(k1_xboole_0,X41)&v5_relat_1(k1_xboole_0,X42)), inference(variable_rename,[status(thm)],[l222_relat_1])).
cnf(c_0_69, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_70, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_55])).
fof(c_0_71, plain, ![X72, X73]:((~v1_tdlat_3(X72)|(~m1_subset_1(X73,k1_zfmisc_1(u1_struct_0(X72)))|v4_pre_topc(X73,X72))|(~v2_pre_topc(X72)|~l1_pre_topc(X72)))&((m1_subset_1(esk9_1(X72),k1_zfmisc_1(u1_struct_0(X72)))|v1_tdlat_3(X72)|(~v2_pre_topc(X72)|~l1_pre_topc(X72)))&(~v4_pre_topc(esk9_1(X72),X72)|v1_tdlat_3(X72)|(~v2_pre_topc(X72)|~l1_pre_topc(X72))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t18_tdlat_3])])])])])).
cnf(c_0_72, plain, (k8_relset_1(X1,X2,X3,X2)=k1_relset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_73, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)=k1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_57, c_0_34])).
cnf(c_0_74, negated_conjecture, (X1=esk3_0|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))|~v1_xboole_0(X1)|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_75, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_76, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_77, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|~l1_struct_0(esk2_0)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_78, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_79, plain, (m1_subset_1(esk8_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|v5_pre_topc(X3,X1,X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_80, negated_conjecture, (v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_81, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_82, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_83, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_66])).
fof(c_0_84, plain, ![X37]:(~r1_tarski(X37,k1_xboole_0)|X37=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_xboole_1])])).
cnf(c_0_85, plain, (r1_tarski(k1_relat_1(X1),X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_86, plain, (v4_relat_1(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_87, plain, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_88, plain, (v4_pre_topc(X2,X1)|~v1_tdlat_3(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_89, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,u1_struct_0(esk2_0))=k1_relat_1(esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_34]), c_0_73])).
cnf(c_0_90, negated_conjecture, (k1_xboole_0=esk3_0|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76])])).
cnf(c_0_91, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_51])])).
cnf(c_0_92, negated_conjecture, (m1_subset_1(esk8_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_49]), c_0_34]), c_0_50]), c_0_51]), c_0_40]), c_0_80]), c_0_81])]), c_0_52])).
cnf(c_0_93, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_94, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_95, plain, (r1_tarski(k1_relat_1(k1_xboole_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_87])])).
fof(c_0_96, plain, ![X59, X60]:((~v4_pre_topc(X60,X59)|k2_pre_topc(X59,X60)=X60|~m1_subset_1(X60,k1_zfmisc_1(u1_struct_0(X59)))|~l1_pre_topc(X59))&(~v2_pre_topc(X59)|k2_pre_topc(X59,X60)!=X60|v4_pre_topc(X60,X59)|~m1_subset_1(X60,k1_zfmisc_1(u1_struct_0(X59)))|~l1_pre_topc(X59))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t52_pre_topc])])])])).
cnf(c_0_97, plain, (v4_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v1_tdlat_3(X2)), inference(csr,[status(thm)],[c_0_88, c_0_32])).
cnf(c_0_98, negated_conjecture, (m1_subset_1(k1_relat_1(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_39, c_0_89])).
fof(c_0_99, plain, ![X49, X50, X51]:(~v1_xboole_0(X49)|(~m1_subset_1(X51,k1_zfmisc_1(k2_zfmisc_1(X50,X49)))|v1_xboole_0(X51))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])).
cnf(c_0_100, negated_conjecture, (k1_xboole_0=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_91]), c_0_70]), c_0_76])])).
cnf(c_0_101, negated_conjecture, (r1_tarski(esk8_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_36, c_0_92])).
cnf(c_0_102, plain, (k1_relset_1(X1,X2,k2_zfmisc_1(X1,X2))=k1_relat_1(k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_57, c_0_93])).
cnf(c_0_103, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_104, plain, (k2_pre_topc(X2,X1)=X1|~v4_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_105, negated_conjecture, (v4_pre_topc(k1_relat_1(esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_40]), c_0_41])])).
cnf(c_0_106, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_107, plain, (v5_pre_topc(X3,X1,X2)|~r1_tarski(k2_pre_topc(X1,k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk8_3(X1,X2,X3))),k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,k2_pre_topc(X2,esk8_3(X1,X2,X3))))|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_108, negated_conjecture, (u1_struct_0(esk2_0)=esk3_0), inference(rw,[status(thm)],[c_0_91, c_0_100])).
cnf(c_0_109, plain, (k2_zfmisc_1(X1,esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_100]), c_0_100])).
cnf(c_0_110, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),X1)|~v1_xboole_0(u1_struct_0(esk2_0))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_49, c_0_59])).
cnf(c_0_111, negated_conjecture, (esk8_3(esk1_0,esk2_0,esk3_0)=u1_struct_0(esk2_0)|~r1_tarski(u1_struct_0(esk2_0),esk8_3(esk1_0,esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_44, c_0_101])).
cnf(c_0_112, plain, (k8_relset_1(X1,X2,k2_zfmisc_1(X1,X2),X2)=k1_relat_1(k2_zfmisc_1(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_93]), c_0_102])).
cnf(c_0_113, plain, (k1_relat_1(esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_100]), c_0_100])).
cnf(c_0_114, negated_conjecture, (k2_pre_topc(esk1_0,k1_relat_1(esk3_0))=k1_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_98]), c_0_40])]), c_0_105])])).
cnf(c_0_115, negated_conjecture, (v1_xboole_0(esk3_0)|~v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_106, c_0_34])).
cnf(c_0_116, negated_conjecture, (v5_pre_topc(X1,X2,esk2_0)|~r1_tarski(k2_pre_topc(X2,k8_relset_1(u1_struct_0(X2),esk3_0,X1,esk8_3(X2,esk2_0,X1))),k8_relset_1(u1_struct_0(X2),esk3_0,X1,k2_pre_topc(esk2_0,esk8_3(X2,esk2_0,X1))))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0))|~v1_funct_2(X1,u1_struct_0(X2),esk3_0)|~v1_funct_1(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_51]), c_0_80])]), c_0_109])).
cnf(c_0_117, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),X1)|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_110, c_0_91]), c_0_76])])).
cnf(c_0_118, negated_conjecture, (esk8_3(esk1_0,esk2_0,esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_111, c_0_91]), c_0_91]), c_0_75])]), c_0_100])).
cnf(c_0_119, plain, (k8_relset_1(X1,esk3_0,esk3_0,esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_109]), c_0_113])).
cnf(c_0_120, negated_conjecture, (k2_pre_topc(esk1_0,esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_114, c_0_113]), c_0_113])).
cnf(c_0_121, plain, (r1_tarski(esk3_0,X1)), inference(rw,[status(thm)],[c_0_75, c_0_100])).
cnf(c_0_122, negated_conjecture, (v1_xboole_0(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_115, c_0_91]), c_0_76])])).
cnf(c_0_123, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_118]), c_0_119]), c_0_120]), c_0_118]), c_0_121]), c_0_93]), c_0_50]), c_0_40]), c_0_81]), c_0_122])]), c_0_52]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 124
# Proof object clause steps            : 74
# Proof object formula steps           : 50
# Proof object conjectures             : 39
# Proof object clause conjectures      : 36
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 34
# Proof object initial formulas used   : 25
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 78
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 30
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 64
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 63
# Processed clauses                    : 2907
# ...of these trivial                  : 59
# ...subsumed                          : 2180
# ...remaining for further processing  : 668
# Other redundant clauses eliminated   : 33
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 90
# Backward-rewritten                   : 154
# Generated clauses                    : 10447
# ...of the previous two non-trivial   : 9654
# Contextual simplify-reflections      : 24
# Paramodulations                      : 10412
# Factorizations                       : 2
# Equation resolutions                 : 33
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 418
#    Positive orientable unit clauses  : 54
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 361
# Current number of unprocessed clauses: 6062
# ...number of literals in the above   : 32791
# Current number of archived formulas  : 0
# Current number of archived clauses   : 245
# Clause-clause subsumption calls (NU) : 93086
# Rec. Clause-clause subsumption calls : 33122
# Non-unit clause-clause subsumptions  : 1837
# Unit Clause-clause subsumption calls : 1805
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 102
# BW rewrite match successes           : 22
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 146503

# -------------------------------------------------
# User time                : 0.131 s
# System time              : 0.006 s
# Total time               : 0.138 s
# Maximum resident set size: 28752 pages
