# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/613f6cde03b8414abc1fd77f8ef0da0f.E_input', d5_xboole_0)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/613f6cde03b8414abc1fd77f8ef0da0f.E_input', t3_boole)).
fof(t7_tarski, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&r2_hidden(X4,X3))))))), file('data/MPTP/proofs/613f6cde03b8414abc1fd77f8ef0da0f.E_input', t7_tarski)).
fof(d12_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r2_waybel_0(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X4,X5)&r2_hidden(k2_waybel_0(X1,X2,X5),X3))))))), file('data/MPTP/proofs/613f6cde03b8414abc1fd77f8ef0da0f.E_input', d12_waybel_0)).
fof(t10_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r2_waybel_0(X1,X2,X3)<=>~(r1_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3)))))), file('data/MPTP/proofs/613f6cde03b8414abc1fd77f8ef0da0f.E_input', t10_waybel_0)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/613f6cde03b8414abc1fd77f8ef0da0f.E_input', redefinition_k6_subset_1)).
fof(existence_m1_subset_1, axiom, ![X1]:?[X2]:m1_subset_1(X2,X1), file('data/MPTP/proofs/613f6cde03b8414abc1fd77f8ef0da0f.E_input', existence_m1_subset_1)).
fof(t29_yellow_6, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>r1_waybel_0(X1,X2,u1_struct_0(X1)))), file('data/MPTP/proofs/613f6cde03b8414abc1fd77f8ef0da0f.E_input', t29_yellow_6)).
fof(c_0_8, plain, ![X6, X7, X8, X9, X10, X11, X12, X13]:((((r2_hidden(X9,X6)|~r2_hidden(X9,X8)|X8!=k4_xboole_0(X6,X7))&(~r2_hidden(X9,X7)|~r2_hidden(X9,X8)|X8!=k4_xboole_0(X6,X7)))&(~r2_hidden(X10,X6)|r2_hidden(X10,X7)|r2_hidden(X10,X8)|X8!=k4_xboole_0(X6,X7)))&((~r2_hidden(esk1_3(X11,X12,X13),X13)|(~r2_hidden(esk1_3(X11,X12,X13),X11)|r2_hidden(esk1_3(X11,X12,X13),X12))|X13=k4_xboole_0(X11,X12))&((r2_hidden(esk1_3(X11,X12,X13),X11)|r2_hidden(esk1_3(X11,X12,X13),X13)|X13=k4_xboole_0(X11,X12))&(~r2_hidden(esk1_3(X11,X12,X13),X12)|r2_hidden(esk1_3(X11,X12,X13),X13)|X13=k4_xboole_0(X11,X12))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_9, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_10, plain, ![X17]:k4_xboole_0(X17,k1_xboole_0)=X17, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_11, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_9])).
cnf(c_0_12, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_13, plain, ![X18, X19, X21]:((r2_hidden(esk2_2(X18,X19),X19)|~r2_hidden(X18,X19))&(~r2_hidden(X21,X19)|~r2_hidden(X21,esk2_2(X18,X19))|~r2_hidden(X18,X19))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tarski])])])])])).
cnf(c_0_14, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_15, plain, (r2_hidden(esk2_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_16, plain, (~r2_hidden(esk2_2(X1,k1_xboole_0),X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
fof(c_0_17, plain, ![X22, X23, X24, X25, X27, X29]:(((m1_subset_1(esk3_4(X22,X23,X24,X25),u1_struct_0(X23))|~m1_subset_1(X25,u1_struct_0(X23))|~r2_waybel_0(X22,X23,X24)|(v2_struct_0(X23)|~l1_waybel_0(X23,X22))|(v2_struct_0(X22)|~l1_struct_0(X22)))&((r1_orders_2(X23,X25,esk3_4(X22,X23,X24,X25))|~m1_subset_1(X25,u1_struct_0(X23))|~r2_waybel_0(X22,X23,X24)|(v2_struct_0(X23)|~l1_waybel_0(X23,X22))|(v2_struct_0(X22)|~l1_struct_0(X22)))&(r2_hidden(k2_waybel_0(X22,X23,esk3_4(X22,X23,X24,X25)),X24)|~m1_subset_1(X25,u1_struct_0(X23))|~r2_waybel_0(X22,X23,X24)|(v2_struct_0(X23)|~l1_waybel_0(X23,X22))|(v2_struct_0(X22)|~l1_struct_0(X22)))))&((m1_subset_1(esk4_3(X22,X23,X27),u1_struct_0(X23))|r2_waybel_0(X22,X23,X27)|(v2_struct_0(X23)|~l1_waybel_0(X23,X22))|(v2_struct_0(X22)|~l1_struct_0(X22)))&(~m1_subset_1(X29,u1_struct_0(X23))|(~r1_orders_2(X23,esk4_3(X22,X23,X27),X29)|~r2_hidden(k2_waybel_0(X22,X23,X29),X27))|r2_waybel_0(X22,X23,X27)|(v2_struct_0(X23)|~l1_waybel_0(X23,X22))|(v2_struct_0(X22)|~l1_struct_0(X22))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d12_waybel_0])])])])])])])).
fof(c_0_18, plain, ![X34, X35, X36]:((~r2_waybel_0(X34,X35,X36)|~r1_waybel_0(X34,X35,k6_subset_1(u1_struct_0(X34),X36))|(v2_struct_0(X35)|~l1_waybel_0(X35,X34))|(v2_struct_0(X34)|~l1_struct_0(X34)))&(r1_waybel_0(X34,X35,k6_subset_1(u1_struct_0(X34),X36))|r2_waybel_0(X34,X35,X36)|(v2_struct_0(X35)|~l1_waybel_0(X35,X34))|(v2_struct_0(X34)|~l1_struct_0(X34)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t10_waybel_0])])])])])).
fof(c_0_19, plain, ![X15, X16]:k6_subset_1(X15,X16)=k4_xboole_0(X15,X16), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
cnf(c_0_20, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_16, c_0_15])).
cnf(c_0_21, plain, (r2_hidden(k2_waybel_0(X1,X2,esk3_4(X1,X2,X3,X4)),X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~r2_waybel_0(X1,X2,X3)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_22, plain, ![X30]:m1_subset_1(esk5_1(X30),X30), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[existence_m1_subset_1])])).
fof(c_0_23, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>r1_waybel_0(X1,X2,u1_struct_0(X1))))), inference(assume_negation,[status(cth)],[t29_yellow_6])).
cnf(c_0_24, plain, (r1_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3))|r2_waybel_0(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_waybel_0(X1,X2,k1_xboole_0)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27, plain, (m1_subset_1(esk5_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_28, negated_conjecture, ((~v2_struct_0(esk6_0)&l1_struct_0(esk6_0))&((~v2_struct_0(esk7_0)&(v4_orders_2(esk7_0)&(v7_waybel_0(esk7_0)&l1_waybel_0(esk7_0,esk6_0))))&~r1_waybel_0(esk6_0,esk7_0,u1_struct_0(esk6_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_23])])])])).
cnf(c_0_29, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r2_waybel_0(X1,X2,X3)|r1_waybel_0(X1,X2,k4_xboole_0(u1_struct_0(X1),X3))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)), inference(rw,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_30, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X2,X1,k1_xboole_0)|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_31, negated_conjecture, (~r1_waybel_0(esk6_0,esk7_0,u1_struct_0(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_32, plain, (r1_waybel_0(X1,X2,u1_struct_0(X1))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_12]), c_0_30])).
cnf(c_0_33, negated_conjecture, (l1_waybel_0(esk7_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34, negated_conjecture, (l1_struct_0(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, negated_conjecture, (~v2_struct_0(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36, negated_conjecture, (~v2_struct_0(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33]), c_0_34])]), c_0_35]), c_0_36]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 38
# Proof object clause steps            : 21
# Proof object formula steps           : 17
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 24
# Processed clauses                    : 170
# ...of these trivial                  : 2
# ...subsumed                          : 83
# ...remaining for further processing  : 85
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 10
# Generated clauses                    : 221
# ...of the previous two non-trivial   : 204
# Contextual simplify-reflections      : 1
# Paramodulations                      : 216
# Factorizations                       : 2
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 45
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 33
# Current number of unprocessed clauses: 68
# ...number of literals in the above   : 214
# Current number of archived formulas  : 0
# Current number of archived clauses   : 38
# Clause-clause subsumption calls (NU) : 423
# Rec. Clause-clause subsumption calls : 129
# Non-unit clause-clause subsumptions  : 13
# Unit Clause-clause subsumption calls : 9
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 15
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4403

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.003 s
# Total time               : 0.011 s
# Maximum resident set size: 28760 pages
