# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d3_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(r1_tarski(X1,X2)<=>![X3, X4]:(r2_hidden(k4_tarski(X3,X4),X1)=>r2_hidden(k4_tarski(X3,X4),X2)))), file('data/MPTP/proofs/6169bf6951894abe92b8b28e28e509d5.E_input', d3_relat_1)).
fof(dt_k7_relat_1, axiom, ![X1, X2]:(v1_relat_1(X1)=>v1_relat_1(k7_relat_1(X1,X2))), file('data/MPTP/proofs/6169bf6951894abe92b8b28e28e509d5.E_input', dt_k7_relat_1)).
fof(d11_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(v1_relat_1(X3)=>(X3=k7_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>(r2_hidden(X4,X2)&r2_hidden(k4_tarski(X4,X5),X1)))))), file('data/MPTP/proofs/6169bf6951894abe92b8b28e28e509d5.E_input', d11_relat_1)).
fof(t88_relat_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k7_relat_1(X2,X1),X2)), file('data/MPTP/proofs/6169bf6951894abe92b8b28e28e509d5.E_input', t88_relat_1)).
fof(c_0_4, plain, ![X6, X7, X8, X9, X10]:((~r1_tarski(X6,X7)|(~r2_hidden(k4_tarski(X8,X9),X6)|r2_hidden(k4_tarski(X8,X9),X7))|~v1_relat_1(X6))&((r2_hidden(k4_tarski(esk1_2(X6,X10),esk2_2(X6,X10)),X6)|r1_tarski(X6,X10)|~v1_relat_1(X6))&(~r2_hidden(k4_tarski(esk1_2(X6,X10),esk2_2(X6,X10)),X10)|r1_tarski(X6,X10)|~v1_relat_1(X6)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_relat_1])])])])])])).
fof(c_0_5, plain, ![X13, X14]:(~v1_relat_1(X13)|v1_relat_1(k7_relat_1(X13,X14))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_relat_1])])).
fof(c_0_6, plain, ![X15, X16, X17, X18, X19, X20, X21]:((((r2_hidden(X18,X16)|~r2_hidden(k4_tarski(X18,X19),X17)|X17!=k7_relat_1(X15,X16)|~v1_relat_1(X17)|~v1_relat_1(X15))&(r2_hidden(k4_tarski(X18,X19),X15)|~r2_hidden(k4_tarski(X18,X19),X17)|X17!=k7_relat_1(X15,X16)|~v1_relat_1(X17)|~v1_relat_1(X15)))&(~r2_hidden(X20,X16)|~r2_hidden(k4_tarski(X20,X21),X15)|r2_hidden(k4_tarski(X20,X21),X17)|X17!=k7_relat_1(X15,X16)|~v1_relat_1(X17)|~v1_relat_1(X15)))&((~r2_hidden(k4_tarski(esk3_3(X15,X16,X17),esk4_3(X15,X16,X17)),X17)|(~r2_hidden(esk3_3(X15,X16,X17),X16)|~r2_hidden(k4_tarski(esk3_3(X15,X16,X17),esk4_3(X15,X16,X17)),X15))|X17=k7_relat_1(X15,X16)|~v1_relat_1(X17)|~v1_relat_1(X15))&((r2_hidden(esk3_3(X15,X16,X17),X16)|r2_hidden(k4_tarski(esk3_3(X15,X16,X17),esk4_3(X15,X16,X17)),X17)|X17=k7_relat_1(X15,X16)|~v1_relat_1(X17)|~v1_relat_1(X15))&(r2_hidden(k4_tarski(esk3_3(X15,X16,X17),esk4_3(X15,X16,X17)),X15)|r2_hidden(k4_tarski(esk3_3(X15,X16,X17),esk4_3(X15,X16,X17)),X17)|X17=k7_relat_1(X15,X16)|~v1_relat_1(X17)|~v1_relat_1(X15))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_relat_1])])])])])])).
cnf(c_0_7, plain, (r2_hidden(k4_tarski(esk1_2(X1,X2),esk2_2(X1,X2)),X1)|r1_tarski(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, plain, (v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(k4_tarski(X1,X2),X4)|X4!=k7_relat_1(X3,X5)|~v1_relat_1(X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (r2_hidden(k4_tarski(esk1_2(k7_relat_1(X1,X2),X3),esk2_2(k7_relat_1(X1,X2),X3)),k7_relat_1(X1,X2))|r1_tarski(k7_relat_1(X1,X2),X3)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
fof(c_0_11, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k7_relat_1(X2,X1),X2))), inference(assume_negation,[status(cth)],[t88_relat_1])).
cnf(c_0_12, plain, (r2_hidden(k4_tarski(esk1_2(k7_relat_1(X1,X2),X3),esk2_2(k7_relat_1(X1,X2),X3)),X4)|r1_tarski(k7_relat_1(X1,X2),X3)|k7_relat_1(X1,X2)!=k7_relat_1(X4,X5)|~v1_relat_1(X4)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_8])).
fof(c_0_13, negated_conjecture, (v1_relat_1(esk6_0)&~r1_tarski(k7_relat_1(esk6_0,esk5_0),esk6_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
cnf(c_0_14, plain, (r1_tarski(X1,X2)|~r2_hidden(k4_tarski(esk1_2(X1,X2),esk2_2(X1,X2)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_15, plain, (r2_hidden(k4_tarski(esk1_2(k7_relat_1(X1,X2),X3),esk2_2(k7_relat_1(X1,X2),X3)),X1)|r1_tarski(k7_relat_1(X1,X2),X3)|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_12])).
cnf(c_0_16, negated_conjecture, (~r1_tarski(k7_relat_1(esk6_0,esk5_0),esk6_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, plain, (r1_tarski(k7_relat_1(X1,X2),X1)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_8])).
cnf(c_0_18, negated_conjecture, (v1_relat_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 20
# Proof object clause steps            : 11
# Proof object formula steps           : 9
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 6
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 4
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 12
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 12
# Processed clauses                    : 131
# ...of these trivial                  : 0
# ...subsumed                          : 32
# ...remaining for further processing  : 99
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 11
# Backward-rewritten                   : 4
# Generated clauses                    : 303
# ...of the previous two non-trivial   : 277
# Contextual simplify-reflections      : 12
# Paramodulations                      : 294
# Factorizations                       : 0
# Equation resolutions                 : 9
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 72
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 67
# Current number of unprocessed clauses: 159
# ...number of literals in the above   : 741
# Current number of archived formulas  : 0
# Current number of archived clauses   : 27
# Clause-clause subsumption calls (NU) : 945
# Rec. Clause-clause subsumption calls : 248
# Non-unit clause-clause subsumptions  : 29
# Unit Clause-clause subsumption calls : 27
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8351

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.000 s
# Total time               : 0.015 s
# Maximum resident set size: 28632 pages
