# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_waybel_7, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v1_subset_1(X2,u1_struct_0(X1))<=>~(r2_hidden(k3_yellow_0(X1),X2))))), file('data/MPTP/proofs/617ce5e7e3494afe837594630041c1f4.E_input', t8_waybel_7)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/617ce5e7e3494afe837594630041c1f4.E_input', d2_subset_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/617ce5e7e3494afe837594630041c1f4.E_input', t2_subset)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/617ce5e7e3494afe837594630041c1f4.E_input', d1_zfmisc_1)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/617ce5e7e3494afe837594630041c1f4.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/617ce5e7e3494afe837594630041c1f4.E_input', d4_subset_1)).
fof(t8_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,X2)<=>r2_hidden(X4,X3)))=>X2=X3))), file('data/MPTP/proofs/617ce5e7e3494afe837594630041c1f4.E_input', t8_subset_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/617ce5e7e3494afe837594630041c1f4.E_input', d3_tarski)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/617ce5e7e3494afe837594630041c1f4.E_input', d1_xboole_0)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/617ce5e7e3494afe837594630041c1f4.E_input', dt_k1_yellow_0)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/617ce5e7e3494afe837594630041c1f4.E_input', d9_lattice3)).
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/617ce5e7e3494afe837594630041c1f4.E_input', d9_yellow_0)).
fof(d20_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v13_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r1_orders_2(X1,X3,X4))=>r2_hidden(X4,X2))))))), file('data/MPTP/proofs/617ce5e7e3494afe837594630041c1f4.E_input', d20_waybel_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/617ce5e7e3494afe837594630041c1f4.E_input', t4_subset)).
fof(t42_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))=>(r1_yellow_0(X1,k1_xboole_0)&r2_yellow_0(X1,u1_struct_0(X1)))), file('data/MPTP/proofs/617ce5e7e3494afe837594630041c1f4.E_input', t42_yellow_0)).
fof(d7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(v1_subset_1(X2,X1)<=>~(X2=X1))), file('data/MPTP/proofs/617ce5e7e3494afe837594630041c1f4.E_input', d7_subset_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/617ce5e7e3494afe837594630041c1f4.E_input', fc1_xboole_0)).
fof(d11_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0)), file('data/MPTP/proofs/617ce5e7e3494afe837594630041c1f4.E_input', d11_yellow_0)).
fof(c_0_18, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v1_subset_1(X2,u1_struct_0(X1))<=>~(r2_hidden(k3_yellow_0(X1),X2)))))), inference(assume_negation,[status(cth)],[t8_waybel_7])).
fof(c_0_19, plain, ![X40, X41]:(((~m1_subset_1(X41,X40)|r2_hidden(X41,X40)|v1_xboole_0(X40))&(~r2_hidden(X41,X40)|m1_subset_1(X41,X40)|v1_xboole_0(X40)))&((~m1_subset_1(X41,X40)|v1_xboole_0(X41)|~v1_xboole_0(X40))&(~v1_xboole_0(X41)|m1_subset_1(X41,X40)|~v1_xboole_0(X40)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_20, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_yellow_0(esk1_0)&l1_orders_2(esk1_0))))))&((~v1_xboole_0(esk2_0)&(v2_waybel_0(esk2_0,esk1_0)&(v13_waybel_0(esk2_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))))&((~v1_subset_1(esk2_0,u1_struct_0(esk1_0))|r2_hidden(k3_yellow_0(esk1_0),esk2_0))&(v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k3_yellow_0(esk1_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_18])])])])).
fof(c_0_21, plain, ![X21, X22]:(~m1_subset_1(X21,X22)|(v1_xboole_0(X22)|r2_hidden(X21,X22))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_22, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_25, plain, ![X12, X13, X14, X15, X16, X17]:(((~r2_hidden(X14,X13)|r1_tarski(X14,X12)|X13!=k1_zfmisc_1(X12))&(~r1_tarski(X15,X12)|r2_hidden(X15,X13)|X13!=k1_zfmisc_1(X12)))&((~r2_hidden(esk4_2(X16,X17),X17)|~r1_tarski(esk4_2(X16,X17),X16)|X17=k1_zfmisc_1(X16))&(r2_hidden(esk4_2(X16,X17),X17)|r1_tarski(esk4_2(X16,X17),X16)|X17=k1_zfmisc_1(X16)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
cnf(c_0_26, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, negated_conjecture, (~v1_xboole_0(k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24])).
fof(c_0_28, plain, ![X50]:m1_subset_1(k2_subset_1(X50),k1_zfmisc_1(X50)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_29, plain, ![X39]:k2_subset_1(X39)=X39, inference(variable_rename,[status(thm)],[d4_subset_1])).
cnf(c_0_30, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, negated_conjecture, (r2_hidden(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_23]), c_0_27])).
fof(c_0_32, plain, ![X23, X24, X25]:((m1_subset_1(esk5_3(X23,X24,X25),X23)|X24=X25|~m1_subset_1(X25,k1_zfmisc_1(X23))|~m1_subset_1(X24,k1_zfmisc_1(X23)))&((~r2_hidden(esk5_3(X23,X24,X25),X24)|~r2_hidden(esk5_3(X23,X24,X25),X25)|X24=X25|~m1_subset_1(X25,k1_zfmisc_1(X23))|~m1_subset_1(X24,k1_zfmisc_1(X23)))&(r2_hidden(esk5_3(X23,X24,X25),X24)|r2_hidden(esk5_3(X23,X24,X25),X25)|X24=X25|~m1_subset_1(X25,k1_zfmisc_1(X23))|~m1_subset_1(X24,k1_zfmisc_1(X23))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_subset_1])])])])])).
cnf(c_0_33, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_35, plain, ![X53, X54, X55, X56, X57]:((~r1_tarski(X53,X54)|(~r2_hidden(X55,X53)|r2_hidden(X55,X54)))&((r2_hidden(esk10_2(X56,X57),X56)|r1_tarski(X56,X57))&(~r2_hidden(esk10_2(X56,X57),X57)|r1_tarski(X56,X57)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_36, negated_conjecture, (r1_tarski(esk2_0,X1)|k1_zfmisc_1(u1_struct_0(esk1_0))!=k1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_37, plain, (m1_subset_1(esk5_3(X1,X2,X3),X1)|X2=X3|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_39, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_40, negated_conjecture, (r1_tarski(esk2_0,u1_struct_0(esk1_0))), inference(er,[status(thm)],[c_0_36])).
fof(c_0_41, plain, ![X32, X33, X34]:((~v1_xboole_0(X32)|~r2_hidden(X33,X32))&(r2_hidden(esk7_1(X34),X34)|v1_xboole_0(X34))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_42, plain, ![X19, X20]:(~l1_orders_2(X19)|m1_subset_1(k1_yellow_0(X19,X20),u1_struct_0(X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
fof(c_0_43, plain, ![X7, X8, X9, X10]:((~r2_lattice3(X7,X8,X9)|(~m1_subset_1(X10,u1_struct_0(X7))|(~r2_hidden(X10,X8)|r1_orders_2(X7,X10,X9)))|~m1_subset_1(X9,u1_struct_0(X7))|~l1_orders_2(X7))&((m1_subset_1(esk3_3(X7,X8,X9),u1_struct_0(X7))|r2_lattice3(X7,X8,X9)|~m1_subset_1(X9,u1_struct_0(X7))|~l1_orders_2(X7))&((r2_hidden(esk3_3(X7,X8,X9),X8)|r2_lattice3(X7,X8,X9)|~m1_subset_1(X9,u1_struct_0(X7))|~l1_orders_2(X7))&(~r1_orders_2(X7,esk3_3(X7,X8,X9),X9)|r2_lattice3(X7,X8,X9)|~m1_subset_1(X9,u1_struct_0(X7))|~l1_orders_2(X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
cnf(c_0_44, plain, (X1=X2|m1_subset_1(esk5_3(X2,X1,X2),X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_45, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_46, plain, (r2_hidden(esk7_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_47, plain, ![X27, X28, X29, X30]:(((r2_lattice3(X27,X28,X29)|X29!=k1_yellow_0(X27,X28)|~r1_yellow_0(X27,X28)|~m1_subset_1(X29,u1_struct_0(X27))|~l1_orders_2(X27))&(~m1_subset_1(X30,u1_struct_0(X27))|(~r2_lattice3(X27,X28,X30)|r1_orders_2(X27,X29,X30))|X29!=k1_yellow_0(X27,X28)|~r1_yellow_0(X27,X28)|~m1_subset_1(X29,u1_struct_0(X27))|~l1_orders_2(X27)))&((m1_subset_1(esk6_3(X27,X28,X29),u1_struct_0(X27))|~r2_lattice3(X27,X28,X29)|X29=k1_yellow_0(X27,X28)|~r1_yellow_0(X27,X28)|~m1_subset_1(X29,u1_struct_0(X27))|~l1_orders_2(X27))&((r2_lattice3(X27,X28,esk6_3(X27,X28,X29))|~r2_lattice3(X27,X28,X29)|X29=k1_yellow_0(X27,X28)|~r1_yellow_0(X27,X28)|~m1_subset_1(X29,u1_struct_0(X27))|~l1_orders_2(X27))&(~r1_orders_2(X27,X29,esk6_3(X27,X28,X29))|~r2_lattice3(X27,X28,X29)|X29=k1_yellow_0(X27,X28)|~r1_yellow_0(X27,X28)|~m1_subset_1(X29,u1_struct_0(X27))|~l1_orders_2(X27))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
cnf(c_0_48, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_49, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_50, plain, (r2_hidden(esk3_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_51, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|m1_subset_1(esk5_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_44, c_0_23])).
fof(c_0_52, plain, ![X44, X45, X46, X47]:((~v13_waybel_0(X45,X44)|(~m1_subset_1(X46,u1_struct_0(X44))|(~m1_subset_1(X47,u1_struct_0(X44))|(~r2_hidden(X46,X45)|~r1_orders_2(X44,X46,X47)|r2_hidden(X47,X45))))|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X44)))|~l1_orders_2(X44))&((m1_subset_1(esk8_2(X44,X45),u1_struct_0(X44))|v13_waybel_0(X45,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X44)))|~l1_orders_2(X44))&((m1_subset_1(esk9_2(X44,X45),u1_struct_0(X44))|v13_waybel_0(X45,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X44)))|~l1_orders_2(X44))&(((r2_hidden(esk8_2(X44,X45),X45)|v13_waybel_0(X45,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X44)))|~l1_orders_2(X44))&(r1_orders_2(X44,esk8_2(X44,X45),esk9_2(X44,X45))|v13_waybel_0(X45,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X44)))|~l1_orders_2(X44)))&(~r2_hidden(esk9_2(X44,X45),X45)|v13_waybel_0(X45,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X44)))|~l1_orders_2(X44)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d20_waybel_0])])])])])).
fof(c_0_53, plain, ![X36, X37, X38]:(~r2_hidden(X36,X37)|~m1_subset_1(X37,k1_zfmisc_1(X38))|m1_subset_1(X36,X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_54, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_55, negated_conjecture, (r2_hidden(esk7_1(esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_24])).
cnf(c_0_56, plain, (r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~r1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_57, negated_conjecture, (m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_58, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_lattice3(esk1_0,X1,esk5_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)))|r2_hidden(esk3_3(esk1_0,X1,esk5_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0))),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_49])])).
fof(c_0_59, plain, ![X42]:((r1_yellow_0(X42,k1_xboole_0)|(v2_struct_0(X42)|(~v5_orders_2(X42)|(~v1_yellow_0(X42)|~l1_orders_2(X42)))))&(r2_yellow_0(X42,u1_struct_0(X42))|(v2_struct_0(X42)|(~v5_orders_2(X42)|(~v1_yellow_0(X42)|~l1_orders_2(X42)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t42_yellow_0])])])])).
fof(c_0_60, plain, ![X51, X52]:((~v1_subset_1(X52,X51)|X52!=X51|~m1_subset_1(X52,k1_zfmisc_1(X51)))&(X52=X51|v1_subset_1(X52,X51)|~m1_subset_1(X52,k1_zfmisc_1(X51)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_subset_1])])])])).
cnf(c_0_61, plain, (r2_hidden(X4,X1)|~v13_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,X1)|~r1_orders_2(X2,X3,X4)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_62, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_63, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_64, negated_conjecture, (r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),X2)|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,X3)|~r1_yellow_0(esk1_0,X3)|~r2_lattice3(esk1_0,X3,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_49])])).
cnf(c_0_65, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_lattice3(esk1_0,X1,esk5_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_54, c_0_58])).
cnf(c_0_66, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_67, plain, (r1_yellow_0(X1,k1_xboole_0)|v2_struct_0(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_68, negated_conjecture, (v1_yellow_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_69, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_70, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_71, plain, (X1=X2|v1_subset_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_60])).
fof(c_0_72, plain, ![X43]:(~l1_orders_2(X43)|k3_yellow_0(X43)=k1_yellow_0(X43,k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_yellow_0])])).
cnf(c_0_73, plain, (r2_hidden(X1,X2)|~r1_orders_2(X3,X4,X1)|~r2_hidden(X4,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~v13_waybel_0(X2,X3)|~l1_orders_2(X3)), inference(csr,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_74, negated_conjecture, (v13_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_75, plain, (X2=X3|~r2_hidden(esk5_3(X1,X2,X3),X2)|~r2_hidden(esk5_3(X1,X2,X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_76, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_hidden(esk5_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_51]), c_0_63])).
cnf(c_0_77, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),esk5_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)))|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,X2)|~r1_yellow_0(esk1_0,X2)|~r2_lattice3(esk1_0,X2,esk5_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_64, c_0_51])).
cnf(c_0_78, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_lattice3(esk1_0,k1_xboole_0,esk5_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_79, negated_conjecture, (r1_yellow_0(esk1_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_49]), c_0_68]), c_0_69])]), c_0_70])).
cnf(c_0_80, negated_conjecture, (r2_hidden(k3_yellow_0(esk1_0),esk2_0)|~v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_81, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_71, c_0_23])).
cnf(c_0_82, plain, (k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_83, negated_conjecture, (r2_hidden(X1,esk2_0)|~r1_orders_2(esk1_0,X2,X1)|~r2_hidden(X2,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_23]), c_0_74]), c_0_49])])).
cnf(c_0_84, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|~r2_hidden(esk5_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_38]), c_0_23])])).
cnf(c_0_85, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),esk5_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)))|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_79])])).
cnf(c_0_86, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_hidden(k3_yellow_0(esk1_0),esk2_0)), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_87, negated_conjecture, (k3_yellow_0(esk1_0)=k1_yellow_0(esk1_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_82, c_0_49])).
cnf(c_0_88, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k3_yellow_0(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_89, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|~r1_orders_2(esk1_0,X1,esk5_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)))|~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_51]), c_0_84])).
cnf(c_0_90, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,k1_xboole_0),esk5_3(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk1_0)))), inference(er,[status(thm)],[c_0_85])).
cnf(c_0_91, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_hidden(k1_yellow_0(esk1_0,k1_xboole_0),esk2_0)), inference(rw,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_92, plain, (~v1_subset_1(X1,X2)|X1!=X2|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_93, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k1_yellow_0(esk1_0,k1_xboole_0),esk2_0)), inference(rw,[status(thm)],[c_0_88, c_0_87])).
cnf(c_0_94, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_91])).
cnf(c_0_95, plain, (~v1_subset_1(X1,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_92]), c_0_38])])).
cnf(c_0_96, negated_conjecture, (r2_hidden(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_57]), c_0_63])).
cnf(c_0_97, negated_conjecture, (~r2_hidden(k1_yellow_0(esk1_0,k1_xboole_0),esk2_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_94]), c_0_95])).
cnf(c_0_98, negated_conjecture, (r2_hidden(k1_yellow_0(esk1_0,X1),esk2_0)), inference(rw,[status(thm)],[c_0_96, c_0_94])).
cnf(c_0_99, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_98])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 100
# Proof object clause steps            : 64
# Proof object formula steps           : 36
# Proof object conjectures             : 43
# Proof object clause conjectures      : 40
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 29
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 35
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 54
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 53
# Processed clauses                    : 7146
# ...of these trivial                  : 14
# ...subsumed                          : 3661
# ...remaining for further processing  : 3471
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 205
# Backward-rewritten                   : 1524
# Generated clauses                    : 70693
# ...of the previous two non-trivial   : 69412
# Contextual simplify-reflections      : 54
# Paramodulations                      : 70447
# Factorizations                       : 12
# Equation resolutions                 : 234
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1689
#    Positive orientable unit clauses  : 40
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 1645
# Current number of unprocessed clauses: 62086
# ...number of literals in the above   : 288503
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1782
# Clause-clause subsumption calls (NU) : 1038127
# Rec. Clause-clause subsumption calls : 337261
# Non-unit clause-clause subsumptions  : 3294
# Unit Clause-clause subsumption calls : 5204
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 100
# BW rewrite match successes           : 19
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2280901

# -------------------------------------------------
# User time                : 1.689 s
# System time              : 0.027 s
# Total time               : 1.716 s
# Maximum resident set size: 65116 pages
