# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/61964e985394434bb3af5f1aca6d05af.E_input', t12_setfam_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/61964e985394434bb3af5f1aca6d05af.E_input', t70_enumset1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/61964e985394434bb3af5f1aca6d05af.E_input', t3_xboole_0)).
fof(t166_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k10_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k2_relat_1(X3))&(r2_hidden(k4_tarski(X1,X4),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/61964e985394434bb3af5f1aca6d05af.E_input', t166_relat_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/61964e985394434bb3af5f1aca6d05af.E_input', t100_xboole_1)).
fof(t173_relat_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>(k10_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k2_relat_1(X2),X1))), file('data/MPTP/proofs/61964e985394434bb3af5f1aca6d05af.E_input', t173_relat_1)).
fof(t64_zfmisc_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k4_xboole_0(X2,k1_tarski(X3)))<=>(r2_hidden(X1,X2)&~(X1=X3))), file('data/MPTP/proofs/61964e985394434bb3af5f1aca6d05af.E_input', t64_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/61964e985394434bb3af5f1aca6d05af.E_input', t69_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/61964e985394434bb3af5f1aca6d05af.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/61964e985394434bb3af5f1aca6d05af.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/61964e985394434bb3af5f1aca6d05af.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/61964e985394434bb3af5f1aca6d05af.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/61964e985394434bb3af5f1aca6d05af.E_input', t75_enumset1)).
fof(t20_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)=>(r2_hidden(X1,k1_relat_1(X3))&r2_hidden(X2,k2_relat_1(X3))))), file('data/MPTP/proofs/61964e985394434bb3af5f1aca6d05af.E_input', t20_relat_1)).
fof(t4_boole, axiom, ![X1]:k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0, file('data/MPTP/proofs/61964e985394434bb3af5f1aca6d05af.E_input', t4_boole)).
fof(t143_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k9_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k1_relat_1(X3))&(r2_hidden(k4_tarski(X4,X1),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/61964e985394434bb3af5f1aca6d05af.E_input', t143_relat_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/61964e985394434bb3af5f1aca6d05af.E_input', t7_xboole_0)).
fof(t146_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k9_relat_1(X1,k1_relat_1(X1))=k2_relat_1(X1)), file('data/MPTP/proofs/61964e985394434bb3af5f1aca6d05af.E_input', t146_relat_1)).
fof(c_0_18, plain, ![X36, X37]:k1_setfam_1(k2_tarski(X36,X37))=k3_xboole_0(X36,X37), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_19, plain, ![X39, X40]:k1_enumset1(X39,X39,X40)=k2_tarski(X39,X40), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_20, plain, ![X60, X61, X63, X64, X65]:(((r2_hidden(esk5_2(X60,X61),X60)|r1_xboole_0(X60,X61))&(r2_hidden(esk5_2(X60,X61),X61)|r1_xboole_0(X60,X61)))&(~r2_hidden(X65,X63)|~r2_hidden(X65,X64)|~r1_xboole_0(X63,X64))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_21, plain, ![X26, X27, X28, X30]:(((r2_hidden(esk3_3(X26,X27,X28),k2_relat_1(X28))|~r2_hidden(X26,k10_relat_1(X28,X27))|~v1_relat_1(X28))&((r2_hidden(k4_tarski(X26,esk3_3(X26,X27,X28)),X28)|~r2_hidden(X26,k10_relat_1(X28,X27))|~v1_relat_1(X28))&(r2_hidden(esk3_3(X26,X27,X28),X27)|~r2_hidden(X26,k10_relat_1(X28,X27))|~v1_relat_1(X28))))&(~r2_hidden(X30,k2_relat_1(X28))|(~r2_hidden(k4_tarski(X26,X30),X28)|~r2_hidden(X30,X27))|r2_hidden(X26,k10_relat_1(X28,X27))|~v1_relat_1(X28))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t166_relat_1])])])])])).
fof(c_0_22, plain, ![X58, X59]:k4_xboole_0(X58,X59)=k5_xboole_0(X58,k3_xboole_0(X58,X59)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_23, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (r2_hidden(esk3_3(X1,X2,X3),X2)|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_27, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>(k10_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k2_relat_1(X2),X1)))), inference(assume_negation,[status(cth)],[t173_relat_1])).
fof(c_0_28, plain, ![X23, X24, X25]:(((r2_hidden(X23,X24)|~r2_hidden(X23,k4_xboole_0(X24,k1_tarski(X25))))&(X23!=X25|~r2_hidden(X23,k4_xboole_0(X24,k1_tarski(X25)))))&(~r2_hidden(X23,X24)|X23=X25|r2_hidden(X23,k4_xboole_0(X24,k1_tarski(X25))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t64_zfmisc_1])])])])).
fof(c_0_29, plain, ![X22]:k2_tarski(X22,X22)=k1_tarski(X22), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_30, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_32, plain, ![X19, X20, X21]:k2_enumset1(X19,X19,X20,X21)=k1_enumset1(X19,X20,X21), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_33, plain, ![X10, X11, X12, X13]:k3_enumset1(X10,X10,X11,X12,X13)=k2_enumset1(X10,X11,X12,X13), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_34, plain, ![X14, X15, X16, X17, X18]:k4_enumset1(X14,X14,X15,X16,X17,X18)=k3_enumset1(X14,X15,X16,X17,X18), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_35, plain, ![X52, X53, X54, X55, X56, X57]:k5_enumset1(X52,X52,X53,X54,X55,X56,X57)=k4_enumset1(X52,X53,X54,X55,X56,X57), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_36, plain, ![X44, X45, X46, X47, X48, X49, X50]:k6_enumset1(X44,X44,X45,X46,X47,X48,X49,X50)=k5_enumset1(X44,X45,X46,X47,X48,X49,X50), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_37, plain, ![X41, X42, X43]:((r2_hidden(X41,k1_relat_1(X43))|~r2_hidden(k4_tarski(X41,X42),X43)|~v1_relat_1(X43))&(r2_hidden(X42,k2_relat_1(X43))|~r2_hidden(k4_tarski(X41,X42),X43)|~v1_relat_1(X43))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t20_relat_1])])])).
cnf(c_0_38, plain, (~r2_hidden(esk3_3(X1,X2,X3),X4)|~r2_hidden(X1,k10_relat_1(X3,X2))|~r1_xboole_0(X4,X2)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_39, plain, (r2_hidden(esk3_3(X1,X2,X3),k2_relat_1(X3))|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_40, negated_conjecture, (v1_relat_1(esk2_0)&((k10_relat_1(esk2_0,esk1_0)!=k1_xboole_0|~r1_xboole_0(k2_relat_1(esk2_0),esk1_0))&(k10_relat_1(esk2_0,esk1_0)=k1_xboole_0|r1_xboole_0(k2_relat_1(esk2_0),esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])])).
cnf(c_0_41, plain, (X1!=X2|~r2_hidden(X1,k4_xboole_0(X3,k1_tarski(X2)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_42, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_43, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_44, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_45, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_46, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_47, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_48, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_49, plain, ![X51]:k4_xboole_0(k1_xboole_0,X51)=k1_xboole_0, inference(variable_rename,[status(thm)],[t4_boole])).
cnf(c_0_50, plain, (r2_hidden(X3,k10_relat_1(X2,X4))|~r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)|~r2_hidden(X1,X4)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_51, plain, (r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_52, plain, ![X31, X32, X33, X35]:(((r2_hidden(esk4_3(X31,X32,X33),k1_relat_1(X33))|~r2_hidden(X31,k9_relat_1(X33,X32))|~v1_relat_1(X33))&((r2_hidden(k4_tarski(esk4_3(X31,X32,X33),X31),X33)|~r2_hidden(X31,k9_relat_1(X33,X32))|~v1_relat_1(X33))&(r2_hidden(esk4_3(X31,X32,X33),X32)|~r2_hidden(X31,k9_relat_1(X33,X32))|~v1_relat_1(X33))))&(~r2_hidden(X35,k1_relat_1(X33))|(~r2_hidden(k4_tarski(X35,X31),X33)|~r2_hidden(X35,X32))|r2_hidden(X31,k9_relat_1(X33,X32))|~v1_relat_1(X33))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t143_relat_1])])])])])).
cnf(c_0_53, plain, (~r2_hidden(X1,k10_relat_1(X2,X3))|~r1_xboole_0(k2_relat_1(X2),X3)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_54, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)=k1_xboole_0|r1_xboole_0(k2_relat_1(esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_55, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_56, plain, ![X66]:(X66=k1_xboole_0|r2_hidden(esk6_1(X66),X66)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_57, plain, (X1!=X2|~r2_hidden(X1,k5_xboole_0(X3,k1_setfam_1(k6_enumset1(X3,X3,X3,X3,X3,X3,X3,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_42]), c_0_24]), c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48])).
cnf(c_0_58, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_59, plain, (r2_hidden(X1,k10_relat_1(X2,X3))|~r2_hidden(k4_tarski(X1,X4),X2)|~r2_hidden(X4,X3)|~v1_relat_1(X2)), inference(csr,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_60, plain, (r2_hidden(k4_tarski(esk4_3(X1,X2,X3),X1),X3)|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_61, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)=k1_xboole_0|~r2_hidden(X1,k10_relat_1(esk2_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55])])).
cnf(c_0_62, plain, (X1=k1_xboole_0|r2_hidden(esk6_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_63, plain, (~r2_hidden(X1,k5_xboole_0(X2,k1_setfam_1(k6_enumset1(X2,X2,X2,X2,X2,X2,X2,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)))))), inference(er,[status(thm)],[c_0_57])).
cnf(c_0_64, plain, (k5_xboole_0(k1_xboole_0,k1_setfam_1(k6_enumset1(k1_xboole_0,k1_xboole_0,k1_xboole_0,k1_xboole_0,k1_xboole_0,k1_xboole_0,k1_xboole_0,X1)))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_43]), c_0_44]), c_0_45]), c_0_46]), c_0_47]), c_0_48])).
cnf(c_0_65, plain, (r2_hidden(esk4_3(X1,X2,X3),k10_relat_1(X3,X4))|~r2_hidden(X1,k9_relat_1(X3,X2))|~r2_hidden(X1,X4)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_66, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_67, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
fof(c_0_68, plain, ![X38]:(~v1_relat_1(X38)|k9_relat_1(X38,k1_relat_1(X38))=k2_relat_1(X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t146_relat_1])])).
cnf(c_0_69, negated_conjecture, (~r2_hidden(X1,k9_relat_1(esk2_0,X2))|~r2_hidden(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_55])]), c_0_67])).
cnf(c_0_70, plain, (k9_relat_1(X1,k1_relat_1(X1))=k2_relat_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_71, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_72, negated_conjecture, (~r2_hidden(X1,k2_relat_1(esk2_0))|~r2_hidden(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_55])])).
cnf(c_0_73, plain, (r2_hidden(esk5_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_74, plain, (r1_xboole_0(X1,X2)|~r2_hidden(esk5_2(X1,X2),X3)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_25, c_0_71])).
cnf(c_0_75, negated_conjecture, (r1_xboole_0(X1,k2_relat_1(esk2_0))|~r2_hidden(esk5_2(X1,k2_relat_1(esk2_0)),esk1_0)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_76, negated_conjecture, (k10_relat_1(esk2_0,esk1_0)!=k1_xboole_0|~r1_xboole_0(k2_relat_1(esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_77, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_74, c_0_73])).
cnf(c_0_78, negated_conjecture, (r1_xboole_0(esk1_0,k2_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_75, c_0_71])).
cnf(c_0_79, negated_conjecture, (~r1_xboole_0(k2_relat_1(esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_66])])).
cnf(c_0_80, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_79]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 81
# Proof object clause steps            : 44
# Proof object formula steps           : 37
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 33
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 22
# Processed clauses                    : 1002
# ...of these trivial                  : 1
# ...subsumed                          : 780
# ...remaining for further processing  : 221
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 4
# Generated clauses                    : 2772
# ...of the previous two non-trivial   : 2615
# Contextual simplify-reflections      : 2
# Paramodulations                      : 2771
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 194
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 185
# Current number of unprocessed clauses: 1653
# ...number of literals in the above   : 6823
# Current number of archived formulas  : 0
# Current number of archived clauses   : 35
# Clause-clause subsumption calls (NU) : 7866
# Rec. Clause-clause subsumption calls : 3501
# Non-unit clause-clause subsumptions  : 505
# Unit Clause-clause subsumption calls : 8
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 51885

# -------------------------------------------------
# User time                : 0.036 s
# System time              : 0.000 s
# Total time               : 0.036 s
# Maximum resident set size: 28792 pages
