# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t33_tops_2, conjecture, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_pre_topc(X3,X1)=>(v3_pre_topc(X2,X1)=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))=>(X4=X2=>v3_pre_topc(X4,X3))))))), file('data/MPTP/proofs/620bcb0e16da4206b34d947731c64d7b.E_input', t33_tops_2)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/620bcb0e16da4206b34d947731c64d7b.E_input', dt_m1_pre_topc)).
fof(d10_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:((v1_pre_topc(X3)&m1_pre_topc(X3,X1))=>(X3=k1_pre_topc(X1,X2)<=>k2_struct_0(X3)=X2)))), file('data/MPTP/proofs/620bcb0e16da4206b34d947731c64d7b.E_input', d10_pre_topc)).
fof(dt_k1_pre_topc, axiom, ![X1, X2]:((l1_pre_topc(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_pre_topc(k1_pre_topc(X1,X2))&m1_pre_topc(k1_pre_topc(X1,X2),X1))), file('data/MPTP/proofs/620bcb0e16da4206b34d947731c64d7b.E_input', dt_k1_pre_topc)).
fof(redefinition_k9_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>k9_subset_1(X1,X2,X3)=k3_xboole_0(X2,X3)), file('data/MPTP/proofs/620bcb0e16da4206b34d947731c64d7b.E_input', redefinition_k9_subset_1)).
fof(d9_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>(m1_pre_topc(X2,X1)<=>(r1_tarski(k2_struct_0(X2),k2_struct_0(X1))&![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>(r2_hidden(X3,u1_pre_topc(X2))<=>?[X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))&(r2_hidden(X4,u1_pre_topc(X1))&X3=k9_subset_1(u1_struct_0(X2),X4,k2_struct_0(X2)))))))))), file('data/MPTP/proofs/620bcb0e16da4206b34d947731c64d7b.E_input', d9_pre_topc)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/620bcb0e16da4206b34d947731c64d7b.E_input', commutativity_k3_xboole_0)).
fof(commutativity_k9_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>k9_subset_1(X1,X2,X3)=k9_subset_1(X1,X3,X2)), file('data/MPTP/proofs/620bcb0e16da4206b34d947731c64d7b.E_input', commutativity_k9_subset_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/620bcb0e16da4206b34d947731c64d7b.E_input', t28_xboole_1)).
fof(t32_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>(v3_pre_topc(X3,X2)<=>?[X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))&(v3_pre_topc(X4,X1)&k9_subset_1(u1_struct_0(X2),X4,k2_struct_0(X2))=X3)))))), file('data/MPTP/proofs/620bcb0e16da4206b34d947731c64d7b.E_input', t32_tops_2)).
fof(c_0_10, negated_conjecture, ~(![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_pre_topc(X3,X1)=>(v3_pre_topc(X2,X1)=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))=>(X4=X2=>v3_pre_topc(X4,X3)))))))), inference(assume_negation,[status(cth)],[t33_tops_2])).
fof(c_0_11, negated_conjecture, (l1_pre_topc(esk1_0)&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_pre_topc(esk3_0,esk1_0)&(v3_pre_topc(esk2_0,esk1_0)&(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk3_0)))&(esk4_0=esk2_0&~v3_pre_topc(esk4_0,esk3_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
fof(c_0_12, plain, ![X13, X14]:(~l1_pre_topc(X13)|(~m1_pre_topc(X14,X13)|l1_pre_topc(X14))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
fof(c_0_13, plain, ![X31, X32, X33]:((X33!=k1_pre_topc(X31,X32)|k2_struct_0(X33)=X32|(~v1_pre_topc(X33)|~m1_pre_topc(X33,X31))|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|~l1_pre_topc(X31))&(k2_struct_0(X33)!=X32|X33=k1_pre_topc(X31,X32)|(~v1_pre_topc(X33)|~m1_pre_topc(X33,X31))|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|~l1_pre_topc(X31))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_pre_topc])])])])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, negated_conjecture, (esk4_0=esk2_0), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_19, plain, ![X9, X10]:((v1_pre_topc(k1_pre_topc(X9,X10))|(~l1_pre_topc(X9)|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))))&(m1_pre_topc(k1_pre_topc(X9,X10),X9)|(~l1_pre_topc(X9)|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_pre_topc])])])).
fof(c_0_20, plain, ![X23, X24, X25]:(~m1_subset_1(X25,k1_zfmisc_1(X23))|k9_subset_1(X23,X24,X25)=k3_xboole_0(X24,X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k9_subset_1])])).
fof(c_0_21, plain, ![X15, X16, X17, X19, X21]:(((r1_tarski(k2_struct_0(X16),k2_struct_0(X15))|~m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))&(((m1_subset_1(esk5_3(X15,X16,X17),k1_zfmisc_1(u1_struct_0(X15)))|~r2_hidden(X17,u1_pre_topc(X16))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))&((r2_hidden(esk5_3(X15,X16,X17),u1_pre_topc(X15))|~r2_hidden(X17,u1_pre_topc(X16))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))&(X17=k9_subset_1(u1_struct_0(X16),esk5_3(X15,X16,X17),k2_struct_0(X16))|~r2_hidden(X17,u1_pre_topc(X16))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))))&(~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X15)))|(~r2_hidden(X19,u1_pre_topc(X15))|X17!=k9_subset_1(u1_struct_0(X16),X19,k2_struct_0(X16)))|r2_hidden(X17,u1_pre_topc(X16))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))))&((m1_subset_1(esk6_2(X15,X16),k1_zfmisc_1(u1_struct_0(X16)))|~r1_tarski(k2_struct_0(X16),k2_struct_0(X15))|m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))&((~r2_hidden(esk6_2(X15,X16),u1_pre_topc(X16))|(~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X15)))|(~r2_hidden(X21,u1_pre_topc(X15))|esk6_2(X15,X16)!=k9_subset_1(u1_struct_0(X16),X21,k2_struct_0(X16))))|~r1_tarski(k2_struct_0(X16),k2_struct_0(X15))|m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))&((m1_subset_1(esk7_2(X15,X16),k1_zfmisc_1(u1_struct_0(X15)))|r2_hidden(esk6_2(X15,X16),u1_pre_topc(X16))|~r1_tarski(k2_struct_0(X16),k2_struct_0(X15))|m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))&((r2_hidden(esk7_2(X15,X16),u1_pre_topc(X15))|r2_hidden(esk6_2(X15,X16),u1_pre_topc(X16))|~r1_tarski(k2_struct_0(X16),k2_struct_0(X15))|m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))&(esk6_2(X15,X16)=k9_subset_1(u1_struct_0(X16),esk7_2(X15,X16),k2_struct_0(X16))|r2_hidden(esk6_2(X15,X16),u1_pre_topc(X16))|~r1_tarski(k2_struct_0(X16),k2_struct_0(X15))|m1_pre_topc(X16,X15)|~l1_pre_topc(X16)|~l1_pre_topc(X15))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_pre_topc])])])])])).
cnf(c_0_22, plain, (k2_struct_0(X1)=X3|X1!=k1_pre_topc(X2,X3)|~v1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk3_0)))), inference(rw,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_24, negated_conjecture, (l1_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])])).
cnf(c_0_25, plain, (v1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (m1_pre_topc(k1_pre_topc(X1,X2),X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (k9_subset_1(X2,X3,X1)=k3_xboole_0(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_29, plain, (r1_tarski(k2_struct_0(X1),k2_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X1)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, negated_conjecture, (k2_struct_0(X1)=esk2_0|X1!=k1_pre_topc(esk3_0,esk2_0)|~v1_pre_topc(X1)|~m1_pre_topc(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24])])).
cnf(c_0_31, negated_conjecture, (v1_pre_topc(k1_pre_topc(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_23]), c_0_24])])).
cnf(c_0_32, negated_conjecture, (m1_pre_topc(k1_pre_topc(esk3_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_23]), c_0_24])])).
fof(c_0_33, plain, ![X29, X30]:k3_xboole_0(X29,X30)=k3_xboole_0(X30,X29), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_34, plain, ![X26, X27, X28]:(~m1_subset_1(X28,k1_zfmisc_1(X26))|k9_subset_1(X26,X27,X28)=k9_subset_1(X26,X28,X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k9_subset_1])])).
cnf(c_0_35, negated_conjecture, (k3_xboole_0(X1,esk2_0)=k9_subset_1(u1_struct_0(esk1_0),X1,esk2_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
fof(c_0_36, plain, ![X11, X12]:(~r1_tarski(X11,X12)|k3_xboole_0(X11,X12)=X11), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_37, plain, (r1_tarski(k2_struct_0(X1),k2_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[c_0_29, c_0_16])).
cnf(c_0_38, negated_conjecture, (k2_struct_0(k1_pre_topc(esk3_0,esk2_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_30]), c_0_31]), c_0_32])])).
cnf(c_0_39, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_40, plain, (k9_subset_1(X2,X3,X1)=k9_subset_1(X2,X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_41, plain, ![X34, X35, X36, X38]:(((m1_subset_1(esk8_3(X34,X35,X36),k1_zfmisc_1(u1_struct_0(X34)))|~v3_pre_topc(X36,X35)|~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X35)))|~m1_pre_topc(X35,X34)|~l1_pre_topc(X34))&((v3_pre_topc(esk8_3(X34,X35,X36),X34)|~v3_pre_topc(X36,X35)|~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X35)))|~m1_pre_topc(X35,X34)|~l1_pre_topc(X34))&(k9_subset_1(u1_struct_0(X35),esk8_3(X34,X35,X36),k2_struct_0(X35))=X36|~v3_pre_topc(X36,X35)|~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X35)))|~m1_pre_topc(X35,X34)|~l1_pre_topc(X34))))&(~m1_subset_1(X38,k1_zfmisc_1(u1_struct_0(X34)))|(~v3_pre_topc(X38,X34)|k9_subset_1(u1_struct_0(X35),X38,k2_struct_0(X35))!=X36)|v3_pre_topc(X36,X35)|~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X35)))|~m1_pre_topc(X35,X34)|~l1_pre_topc(X34))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t32_tops_2])])])])])).
cnf(c_0_42, negated_conjecture, (k9_subset_1(u1_struct_0(esk3_0),X1,esk2_0)=k9_subset_1(u1_struct_0(esk1_0),X1,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_23]), c_0_35])).
cnf(c_0_43, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44, negated_conjecture, (r1_tarski(esk2_0,k2_struct_0(esk3_0))), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_32]), c_0_24])]), c_0_38])).
cnf(c_0_45, negated_conjecture, (k3_xboole_0(esk2_0,X1)=k9_subset_1(u1_struct_0(esk1_0),X1,esk2_0)), inference(spm,[status(thm)],[c_0_39, c_0_35])).
cnf(c_0_46, negated_conjecture, (k9_subset_1(u1_struct_0(esk1_0),esk2_0,X1)=k9_subset_1(u1_struct_0(esk1_0),X1,esk2_0)), inference(spm,[status(thm)],[c_0_40, c_0_28])).
cnf(c_0_47, plain, (v3_pre_topc(X4,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v3_pre_topc(X1,X2)|k9_subset_1(u1_struct_0(X3),X1,k2_struct_0(X3))!=X4|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_48, negated_conjecture, (k9_subset_1(u1_struct_0(esk3_0),esk2_0,X1)=k9_subset_1(u1_struct_0(esk1_0),X1,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_23]), c_0_42])).
cnf(c_0_49, negated_conjecture, (k9_subset_1(u1_struct_0(esk1_0),esk2_0,k2_struct_0(esk3_0))=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_46])).
cnf(c_0_50, negated_conjecture, (~v3_pre_topc(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_51, negated_conjecture, (v3_pre_topc(X1,esk3_0)|esk2_0!=X1|~v3_pre_topc(esk2_0,X2)|~m1_pre_topc(esk3_0,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk3_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_46]), c_0_49])).
cnf(c_0_52, negated_conjecture, (~v3_pre_topc(esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_50, c_0_15])).
cnf(c_0_53, negated_conjecture, (~v3_pre_topc(esk2_0,X1)|~m1_pre_topc(esk3_0,X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_23]), c_0_52])).
cnf(c_0_54, negated_conjecture, (v3_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_55, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_28]), c_0_54]), c_0_17]), c_0_18])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 56
# Proof object clause steps            : 35
# Proof object formula steps           : 21
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 28
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 99
# ...of these trivial                  : 1
# ...subsumed                          : 3
# ...remaining for further processing  : 95
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 74
# ...of the previous two non-trivial   : 65
# Contextual simplify-reflections      : 5
# Paramodulations                      : 70
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 64
#    Positive orientable unit clauses  : 25
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 36
# Current number of unprocessed clauses: 26
# ...number of literals in the above   : 138
# Current number of archived formulas  : 0
# Current number of archived clauses   : 31
# Clause-clause subsumption calls (NU) : 719
# Rec. Clause-clause subsumption calls : 46
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 148
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5058

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.003 s
# Total time               : 0.012 s
# Maximum resident set size: 28640 pages
