# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/624e9057f0fc4cafa35510690ade009f.E_input', d1_tarski)).
fof(d15_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k3_waybel_0(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>?[X5]:(m1_subset_1(X5,u1_struct_0(X1))&(r1_orders_2(X1,X4,X5)&r2_hidden(X5,X2))))))))), file('data/MPTP/proofs/624e9057f0fc4cafa35510690ade009f.E_input', d15_waybel_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/624e9057f0fc4cafa35510690ade009f.E_input', t4_subset)).
fof(dt_k3_waybel_0, axiom, ![X1, X2]:((l1_orders_2(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k3_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/624e9057f0fc4cafa35510690ade009f.E_input', dt_k3_waybel_0)).
fof(t17_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2))))), file('data/MPTP/proofs/624e9057f0fc4cafa35510690ade009f.E_input', t17_waybel_0)).
fof(d17_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>k5_waybel_0(X1,X2)=k3_waybel_0(X1,k6_domain_1(u1_struct_0(X1),X2)))), file('data/MPTP/proofs/624e9057f0fc4cafa35510690ade009f.E_input', d17_waybel_0)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/624e9057f0fc4cafa35510690ade009f.E_input', redefinition_k6_domain_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/624e9057f0fc4cafa35510690ade009f.E_input', t2_subset)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/624e9057f0fc4cafa35510690ade009f.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/624e9057f0fc4cafa35510690ade009f.E_input', dt_l1_orders_2)).
fof(t63_subset_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))), file('data/MPTP/proofs/624e9057f0fc4cafa35510690ade009f.E_input', t63_subset_1)).
fof(c_0_11, plain, ![X6, X7, X8, X9, X10, X11]:(((~r2_hidden(X8,X7)|X8=X6|X7!=k1_tarski(X6))&(X9!=X6|r2_hidden(X9,X7)|X7!=k1_tarski(X6)))&((~r2_hidden(esk1_2(X10,X11),X11)|esk1_2(X10,X11)!=X10|X11=k1_tarski(X10))&(r2_hidden(esk1_2(X10,X11),X11)|esk1_2(X10,X11)=X10|X11=k1_tarski(X10)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_12, plain, ![X13, X14, X15, X16, X18, X20]:((((m1_subset_1(esk2_4(X13,X14,X15,X16),u1_struct_0(X13))|~r2_hidden(X16,X15)|~m1_subset_1(X16,u1_struct_0(X13))|X15!=k3_waybel_0(X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X13)))|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))&((r1_orders_2(X13,X16,esk2_4(X13,X14,X15,X16))|~r2_hidden(X16,X15)|~m1_subset_1(X16,u1_struct_0(X13))|X15!=k3_waybel_0(X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X13)))|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))&(r2_hidden(esk2_4(X13,X14,X15,X16),X14)|~r2_hidden(X16,X15)|~m1_subset_1(X16,u1_struct_0(X13))|X15!=k3_waybel_0(X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X13)))|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))))&(~m1_subset_1(X18,u1_struct_0(X13))|(~r1_orders_2(X13,X16,X18)|~r2_hidden(X18,X14))|r2_hidden(X16,X15)|~m1_subset_1(X16,u1_struct_0(X13))|X15!=k3_waybel_0(X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X13)))|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13)))&((m1_subset_1(esk3_3(X13,X14,X15),u1_struct_0(X13))|X15=k3_waybel_0(X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X13)))|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))&((~r2_hidden(esk3_3(X13,X14,X15),X15)|(~m1_subset_1(X20,u1_struct_0(X13))|(~r1_orders_2(X13,esk3_3(X13,X14,X15),X20)|~r2_hidden(X20,X14)))|X15=k3_waybel_0(X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X13)))|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))&((m1_subset_1(esk4_3(X13,X14,X15),u1_struct_0(X13))|r2_hidden(esk3_3(X13,X14,X15),X15)|X15=k3_waybel_0(X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X13)))|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))&((r1_orders_2(X13,esk3_3(X13,X14,X15),esk4_3(X13,X14,X15))|r2_hidden(esk3_3(X13,X14,X15),X15)|X15=k3_waybel_0(X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X13)))|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))&(r2_hidden(esk4_3(X13,X14,X15),X14)|r2_hidden(esk3_3(X13,X14,X15),X15)|X15=k3_waybel_0(X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X13)))|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~l1_orders_2(X13))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d15_waybel_0])])])])])).
fof(c_0_13, plain, ![X22, X23, X24]:(~r2_hidden(X22,X23)|~m1_subset_1(X23,k1_zfmisc_1(X24))|m1_subset_1(X22,X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_14, plain, ![X25, X26]:(~l1_orders_2(X25)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|m1_subset_1(k3_waybel_0(X25,X26),k1_zfmisc_1(u1_struct_0(X25)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_waybel_0])])).
cnf(c_0_15, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (r2_hidden(esk2_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, plain, (m1_subset_1(k3_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_19, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2)))))), inference(assume_negation,[status(cth)],[t17_waybel_0])).
cnf(c_0_20, plain, (r1_orders_2(X1,X2,esk2_4(X1,X3,X4,X2))|~r2_hidden(X2,X4)|~m1_subset_1(X2,u1_struct_0(X1))|X4!=k3_waybel_0(X1,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, plain, (X1=X2|~r2_hidden(X1,k1_tarski(X2))), inference(er,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (r2_hidden(esk2_4(X1,X2,k3_waybel_0(X1,X2),X3),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~r2_hidden(X3,k3_waybel_0(X1,X2))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_16, c_0_17])]), c_0_18])).
fof(c_0_23, plain, ![X29, X30]:(v2_struct_0(X29)|~l1_orders_2(X29)|(~m1_subset_1(X30,u1_struct_0(X29))|k5_waybel_0(X29,X30)=k3_waybel_0(X29,k6_domain_1(u1_struct_0(X29),X30)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_waybel_0])])])])).
fof(c_0_24, plain, ![X31, X32]:(v1_xboole_0(X31)|~m1_subset_1(X32,X31)|k6_domain_1(X31,X32)=k1_tarski(X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_25, plain, ![X36, X37]:(~m1_subset_1(X36,X37)|(v1_xboole_0(X37)|r2_hidden(X36,X37))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_26, negated_conjecture, ((~v2_struct_0(esk5_0)&l1_orders_2(esk5_0))&(m1_subset_1(esk6_0,u1_struct_0(esk5_0))&(m1_subset_1(esk7_0,u1_struct_0(esk5_0))&((~r2_hidden(esk7_0,k5_waybel_0(esk5_0,esk6_0))|~r1_orders_2(esk5_0,esk7_0,esk6_0))&(r2_hidden(esk7_0,k5_waybel_0(esk5_0,esk6_0))|r1_orders_2(esk5_0,esk7_0,esk6_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_19])])])])).
cnf(c_0_27, plain, (r1_orders_2(X1,X2,esk2_4(X1,X3,k3_waybel_0(X1,X3),X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~r2_hidden(X2,k3_waybel_0(X1,X3))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_20, c_0_17])]), c_0_18])).
cnf(c_0_28, plain, (esk2_4(X1,k1_tarski(X2),k3_waybel_0(X1,k1_tarski(X2)),X3)=X2|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~r2_hidden(X3,k3_waybel_0(X1,k1_tarski(X2)))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_29, plain, (v2_struct_0(X1)|k5_waybel_0(X1,X2)=k3_waybel_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_31, plain, ![X27]:(v2_struct_0(X27)|~l1_struct_0(X27)|~v1_xboole_0(u1_struct_0(X27))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_32, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, negated_conjecture, (m1_subset_1(esk6_0,u1_struct_0(esk5_0))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (r1_orders_2(X1,X2,X3)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~r2_hidden(X2,k3_waybel_0(X1,k1_tarski(X3)))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_35, plain, (k3_waybel_0(X1,k1_tarski(X2))=k5_waybel_0(X1,X2)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_36, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_37, negated_conjecture, (v1_xboole_0(u1_struct_0(esk5_0))|r2_hidden(esk6_0,u1_struct_0(esk5_0))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_38, negated_conjecture, (~v2_struct_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_39, plain, ![X28]:(~l1_orders_2(X28)|l1_struct_0(X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_40, plain, (v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|r1_orders_2(X1,X2,X3)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(X2,k5_waybel_0(X1,X3))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_41, negated_conjecture, (r2_hidden(esk7_0,k5_waybel_0(esk5_0,esk6_0))|r1_orders_2(esk5_0,esk7_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_42, negated_conjecture, (l1_orders_2(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_43, plain, ![X38, X39]:(~r2_hidden(X38,X39)|m1_subset_1(k1_tarski(X38),k1_zfmisc_1(X39))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_subset_1])])).
cnf(c_0_44, negated_conjecture, (r2_hidden(esk6_0,u1_struct_0(esk5_0))|~l1_struct_0(esk5_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_45, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_46, negated_conjecture, (v1_xboole_0(u1_struct_0(esk5_0))|r1_orders_2(esk5_0,esk7_0,esk6_0)|~m1_subset_1(k1_tarski(esk6_0),k1_zfmisc_1(u1_struct_0(esk5_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_33]), c_0_42])]), c_0_38])).
cnf(c_0_47, plain, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_48, negated_conjecture, (r2_hidden(esk6_0,u1_struct_0(esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_42])])).
cnf(c_0_49, negated_conjecture, (v1_xboole_0(u1_struct_0(esk5_0))|r1_orders_2(esk5_0,esk7_0,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])])).
cnf(c_0_50, plain, (m1_subset_1(esk2_4(X1,X2,X3,X4),u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_51, plain, (r2_hidden(X3,X5)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,X3,X1)|~r2_hidden(X1,X4)|~m1_subset_1(X3,u1_struct_0(X2))|X5!=k3_waybel_0(X2,X4)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_52, negated_conjecture, (r1_orders_2(esk5_0,esk7_0,esk6_0)|~l1_struct_0(esk5_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_49]), c_0_38])).
cnf(c_0_53, plain, (m1_subset_1(esk2_4(X1,X2,k3_waybel_0(X1,X2),X3),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~r2_hidden(X3,k3_waybel_0(X1,X2))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_50, c_0_17])]), c_0_18])).
cnf(c_0_54, plain, (r2_hidden(X1,k3_waybel_0(X2,X3))|~r1_orders_2(X2,X1,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~r2_hidden(X4,X3)), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_51, c_0_17])]), c_0_18])).
cnf(c_0_55, negated_conjecture, (r1_orders_2(esk5_0,esk7_0,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_45]), c_0_42])])).
cnf(c_0_56, negated_conjecture, (m1_subset_1(esk7_0,u1_struct_0(esk5_0))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_57, plain, (m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X3,k3_waybel_0(X2,k1_tarski(X1)))), inference(spm,[status(thm)],[c_0_53, c_0_28])).
cnf(c_0_58, negated_conjecture, (r2_hidden(esk7_0,k3_waybel_0(esk5_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk5_0)))|~r2_hidden(esk6_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56]), c_0_42])])).
cnf(c_0_59, negated_conjecture, (~r2_hidden(esk7_0,k5_waybel_0(esk5_0,esk6_0))|~r1_orders_2(esk5_0,esk7_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_60, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk5_0))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk5_0)))|~r2_hidden(esk6_0,k1_tarski(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_42])])).
cnf(c_0_61, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_62, negated_conjecture, (~r2_hidden(esk7_0,k5_waybel_0(esk5_0,esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_55])])).
cnf(c_0_63, negated_conjecture, (v1_xboole_0(u1_struct_0(esk5_0))|r2_hidden(esk7_0,k5_waybel_0(esk5_0,X1))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk5_0)))|~r2_hidden(esk6_0,k1_tarski(X1))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_35]), c_0_42])]), c_0_38]), c_0_60])).
cnf(c_0_64, plain, (r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_61])])).
cnf(c_0_65, negated_conjecture, (v1_xboole_0(u1_struct_0(esk5_0))|~m1_subset_1(k1_tarski(esk6_0),k1_zfmisc_1(u1_struct_0(esk5_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64])])).
cnf(c_0_66, negated_conjecture, (v1_xboole_0(u1_struct_0(esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_47]), c_0_48])])).
cnf(c_0_67, negated_conjecture, (~l1_struct_0(esk5_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_66]), c_0_38])).
cnf(c_0_68, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_45]), c_0_42])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 69
# Proof object clause steps            : 46
# Proof object formula steps           : 23
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 45
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 135
# ...of these trivial                  : 0
# ...subsumed                          : 16
# ...remaining for further processing  : 119
# Other redundant clauses eliminated   : 11
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 13
# Generated clauses                    : 202
# ...of the previous two non-trivial   : 174
# Contextual simplify-reflections      : 12
# Paramodulations                      : 190
# Factorizations                       : 2
# Equation resolutions                 : 11
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 72
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 60
# Current number of unprocessed clauses: 93
# ...number of literals in the above   : 628
# Current number of archived formulas  : 0
# Current number of archived clauses   : 41
# Clause-clause subsumption calls (NU) : 1753
# Rec. Clause-clause subsumption calls : 476
# Non-unit clause-clause subsumptions  : 29
# Unit Clause-clause subsumption calls : 34
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7767

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.003 s
# Total time               : 0.014 s
# Maximum resident set size: 28660 pages
