# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t5_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1)&v2_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1)))), file('data/MPTP/proofs/628e1d5775124a6aa904689d8f16580f.E_input', t5_waybel_0)).
fof(reflexivity_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>r3_orders_2(X1,X2,X2)), file('data/MPTP/proofs/628e1d5775124a6aa904689d8f16580f.E_input', reflexivity_r3_orders_2)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/628e1d5775124a6aa904689d8f16580f.E_input', d1_tarski)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/628e1d5775124a6aa904689d8f16580f.E_input', t69_enumset1)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/628e1d5775124a6aa904689d8f16580f.E_input', redefinition_k6_domain_1)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/628e1d5775124a6aa904689d8f16580f.E_input', redefinition_r3_orders_2)).
fof(d1_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v1_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X3,X2)&(r2_hidden(X4,X2)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>~((r2_hidden(X5,X2)&(r1_orders_2(X1,X3,X5)&r1_orders_2(X1,X4,X5))))))))))))), file('data/MPTP/proofs/628e1d5775124a6aa904689d8f16580f.E_input', d1_waybel_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/628e1d5775124a6aa904689d8f16580f.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/628e1d5775124a6aa904689d8f16580f.E_input', dt_l1_orders_2)).
fof(d2_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v2_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X3,X2)&(r2_hidden(X4,X2)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>~((r2_hidden(X5,X2)&(r1_orders_2(X1,X5,X3)&r1_orders_2(X1,X5,X4))))))))))))), file('data/MPTP/proofs/628e1d5775124a6aa904689d8f16580f.E_input', d2_waybel_0)).
fof(t35_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v6_orders_2(k6_domain_1(u1_struct_0(X1),X2),X1)&m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/628e1d5775124a6aa904689d8f16580f.E_input', t35_orders_2)).
fof(c_0_11, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1)&v2_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1))))), inference(assume_negation,[status(cth)],[t5_waybel_0])).
fof(c_0_12, plain, ![X8, X9, X10]:(v2_struct_0(X8)|(~v3_orders_2(X8)|~l1_orders_2(X8))|(~m1_subset_1(X9,u1_struct_0(X8))|~m1_subset_1(X10,u1_struct_0(X8)))|r3_orders_2(X8,X9,X9)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r3_orders_2])])])).
fof(c_0_13, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&l1_orders_2(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(~v1_waybel_0(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)|~v2_waybel_0(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])).
fof(c_0_14, plain, ![X11, X12, X13, X14, X15, X16]:(((~r2_hidden(X13,X12)|X13=X11|X12!=k1_tarski(X11))&(X14!=X11|r2_hidden(X14,X12)|X12!=k1_tarski(X11)))&((~r2_hidden(esk3_2(X15,X16),X16)|esk3_2(X15,X16)!=X15|X16=k1_tarski(X15))&(r2_hidden(esk3_2(X15,X16),X16)|esk3_2(X15,X16)=X15|X16=k1_tarski(X15)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_15, plain, ![X18]:k2_tarski(X18,X18)=k1_tarski(X18), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_16, plain, ![X19, X20]:(v1_xboole_0(X19)|~m1_subset_1(X20,X19)|k6_domain_1(X19,X20)=k1_tarski(X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_17, plain, ![X21, X22, X23]:((~r3_orders_2(X21,X22,X23)|r1_orders_2(X21,X22,X23)|(v2_struct_0(X21)|(~v3_orders_2(X21)|~l1_orders_2(X21))|(~m1_subset_1(X22,u1_struct_0(X21))|~m1_subset_1(X23,u1_struct_0(X21)))))&(~r1_orders_2(X21,X22,X23)|r3_orders_2(X21,X22,X23)|(v2_struct_0(X21)|(~v3_orders_2(X21)|~l1_orders_2(X21))|(~m1_subset_1(X22,u1_struct_0(X21))|~m1_subset_1(X23,u1_struct_0(X21)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_18, plain, (v2_struct_0(X1)|r3_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_24, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_26, plain, ![X24, X25, X26, X27, X31]:(((m1_subset_1(esk4_4(X24,X25,X26,X27),u1_struct_0(X24))|~r2_hidden(X27,X25)|~r2_hidden(X26,X25)|~m1_subset_1(X27,u1_struct_0(X24))|~m1_subset_1(X26,u1_struct_0(X24))|~v1_waybel_0(X25,X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_orders_2(X24))&((r2_hidden(esk4_4(X24,X25,X26,X27),X25)|~r2_hidden(X27,X25)|~r2_hidden(X26,X25)|~m1_subset_1(X27,u1_struct_0(X24))|~m1_subset_1(X26,u1_struct_0(X24))|~v1_waybel_0(X25,X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_orders_2(X24))&((r1_orders_2(X24,X26,esk4_4(X24,X25,X26,X27))|~r2_hidden(X27,X25)|~r2_hidden(X26,X25)|~m1_subset_1(X27,u1_struct_0(X24))|~m1_subset_1(X26,u1_struct_0(X24))|~v1_waybel_0(X25,X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_orders_2(X24))&(r1_orders_2(X24,X27,esk4_4(X24,X25,X26,X27))|~r2_hidden(X27,X25)|~r2_hidden(X26,X25)|~m1_subset_1(X27,u1_struct_0(X24))|~m1_subset_1(X26,u1_struct_0(X24))|~v1_waybel_0(X25,X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_orders_2(X24)))))&((m1_subset_1(esk5_2(X24,X25),u1_struct_0(X24))|v1_waybel_0(X25,X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_orders_2(X24))&((m1_subset_1(esk6_2(X24,X25),u1_struct_0(X24))|v1_waybel_0(X25,X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_orders_2(X24))&((r2_hidden(esk5_2(X24,X25),X25)|v1_waybel_0(X25,X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_orders_2(X24))&((r2_hidden(esk6_2(X24,X25),X25)|v1_waybel_0(X25,X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_orders_2(X24))&(~m1_subset_1(X31,u1_struct_0(X24))|(~r2_hidden(X31,X25)|(~r1_orders_2(X24,esk5_2(X24,X25),X31)|~r1_orders_2(X24,esk6_2(X24,X25),X31)))|v1_waybel_0(X25,X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_orders_2(X24))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_waybel_0])])])])])).
cnf(c_0_27, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28, negated_conjecture, (r3_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_29, plain, (X1=X3|X2!=k2_tarski(X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_30, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_31, plain, ![X32]:(v2_struct_0(X32)|~l1_struct_0(X32)|~v1_xboole_0(u1_struct_0(X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_32, plain, (k2_tarski(X2,X2)=k6_domain_1(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_25, c_0_24])).
cnf(c_0_33, plain, (v1_waybel_0(X3,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r1_orders_2(X2,esk5_2(X2,X3),X1)|~r1_orders_2(X2,esk6_2(X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_35, plain, (X1=X2|~r2_hidden(X1,k2_tarski(X2,X2))), inference(er,[status(thm)],[c_0_29])).
cnf(c_0_36, plain, (r2_hidden(esk6_2(X1,X2),X2)|v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_37, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X2)), inference(rw,[status(thm)],[c_0_30, c_0_24])).
cnf(c_0_38, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_39, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_32, c_0_19])).
fof(c_0_40, plain, ![X33]:(~l1_orders_2(X33)|l1_struct_0(X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
fof(c_0_41, plain, ![X34, X35, X36, X37, X41]:(((m1_subset_1(esk7_4(X34,X35,X36,X37),u1_struct_0(X34))|~r2_hidden(X37,X35)|~r2_hidden(X36,X35)|~m1_subset_1(X37,u1_struct_0(X34))|~m1_subset_1(X36,u1_struct_0(X34))|~v2_waybel_0(X35,X34)|~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X34)))|~l1_orders_2(X34))&((r2_hidden(esk7_4(X34,X35,X36,X37),X35)|~r2_hidden(X37,X35)|~r2_hidden(X36,X35)|~m1_subset_1(X37,u1_struct_0(X34))|~m1_subset_1(X36,u1_struct_0(X34))|~v2_waybel_0(X35,X34)|~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X34)))|~l1_orders_2(X34))&((r1_orders_2(X34,esk7_4(X34,X35,X36,X37),X36)|~r2_hidden(X37,X35)|~r2_hidden(X36,X35)|~m1_subset_1(X37,u1_struct_0(X34))|~m1_subset_1(X36,u1_struct_0(X34))|~v2_waybel_0(X35,X34)|~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X34)))|~l1_orders_2(X34))&(r1_orders_2(X34,esk7_4(X34,X35,X36,X37),X37)|~r2_hidden(X37,X35)|~r2_hidden(X36,X35)|~m1_subset_1(X37,u1_struct_0(X34))|~m1_subset_1(X36,u1_struct_0(X34))|~v2_waybel_0(X35,X34)|~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X34)))|~l1_orders_2(X34)))))&((m1_subset_1(esk8_2(X34,X35),u1_struct_0(X34))|v2_waybel_0(X35,X34)|~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X34)))|~l1_orders_2(X34))&((m1_subset_1(esk9_2(X34,X35),u1_struct_0(X34))|v2_waybel_0(X35,X34)|~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X34)))|~l1_orders_2(X34))&((r2_hidden(esk8_2(X34,X35),X35)|v2_waybel_0(X35,X34)|~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X34)))|~l1_orders_2(X34))&((r2_hidden(esk9_2(X34,X35),X35)|v2_waybel_0(X35,X34)|~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X34)))|~l1_orders_2(X34))&(~m1_subset_1(X41,u1_struct_0(X34))|(~r2_hidden(X41,X35)|(~r1_orders_2(X34,X41,esk8_2(X34,X35))|~r1_orders_2(X34,X41,esk9_2(X34,X35))))|v2_waybel_0(X35,X34)|~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X34)))|~l1_orders_2(X34))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_waybel_0])])])])])).
cnf(c_0_42, negated_conjecture, (v1_waybel_0(X1,esk1_0)|~r1_orders_2(esk1_0,esk5_2(esk1_0,X1),esk6_2(esk1_0,X1))|~r2_hidden(esk6_2(esk1_0,X1),X1)|~m1_subset_1(esk6_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_20])])).
cnf(c_0_43, plain, (esk6_2(X1,k2_tarski(X2,X2))=X2|v1_waybel_0(k2_tarski(X2,X2),X1)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_44, plain, (r2_hidden(X1,k2_tarski(X1,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_37])])).
cnf(c_0_45, plain, (r2_hidden(esk5_2(X1,X2),X2)|v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_46, plain, ![X42, X43]:((v6_orders_2(k6_domain_1(u1_struct_0(X42),X43),X42)|~m1_subset_1(X43,u1_struct_0(X42))|(v2_struct_0(X42)|(~v3_orders_2(X42)|~l1_orders_2(X42))))&(m1_subset_1(k6_domain_1(u1_struct_0(X42),X43),k1_zfmisc_1(u1_struct_0(X42)))|~m1_subset_1(X43,u1_struct_0(X42))|(v2_struct_0(X42)|(~v3_orders_2(X42)|~l1_orders_2(X42))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t35_orders_2])])])])])).
cnf(c_0_47, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_22])).
cnf(c_0_48, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_49, plain, (v2_waybel_0(X3,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r1_orders_2(X2,X1,esk8_2(X2,X3))|~r1_orders_2(X2,X1,esk9_2(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_50, plain, (r2_hidden(esk9_2(X1,X2),X2)|v2_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_51, negated_conjecture, (v1_waybel_0(k2_tarski(X1,X1),esk1_0)|~r1_orders_2(esk1_0,esk5_2(esk1_0,k2_tarski(X1,X1)),X1)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44]), c_0_20])])).
cnf(c_0_52, plain, (esk5_2(X1,k2_tarski(X2,X2))=X2|v1_waybel_0(k2_tarski(X2,X2),X1)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_35, c_0_45])).
cnf(c_0_53, plain, (m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_54, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_20])])).
cnf(c_0_55, negated_conjecture, (v2_waybel_0(X1,esk1_0)|~r1_orders_2(esk1_0,esk9_2(esk1_0,X1),esk8_2(esk1_0,X1))|~r2_hidden(esk9_2(esk1_0,X1),X1)|~m1_subset_1(esk9_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_34]), c_0_20])])).
cnf(c_0_56, plain, (esk9_2(X1,k2_tarski(X2,X2))=X2|v2_waybel_0(k2_tarski(X2,X2),X1)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_35, c_0_50])).
cnf(c_0_57, plain, (r2_hidden(esk8_2(X1,X2),X2)|v2_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_58, negated_conjecture, (~v1_waybel_0(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)|~v2_waybel_0(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_59, negated_conjecture, (v1_waybel_0(k2_tarski(X1,X1),esk1_0)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_20])]), c_0_34])).
cnf(c_0_60, negated_conjecture, (m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_19]), c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_61, negated_conjecture, (v2_waybel_0(k2_tarski(X1,X1),esk1_0)|~r1_orders_2(esk1_0,X1,esk8_2(esk1_0,k2_tarski(X1,X1)))|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_44]), c_0_20])])).
cnf(c_0_62, plain, (esk8_2(X1,k2_tarski(X2,X2))=X2|v2_waybel_0(k2_tarski(X2,X2),X1)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_35, c_0_57])).
cnf(c_0_63, negated_conjecture, (~v2_waybel_0(k2_tarski(esk2_0,esk2_0),esk1_0)|~v1_waybel_0(k2_tarski(esk2_0,esk2_0),esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_54]), c_0_54])).
cnf(c_0_64, negated_conjecture, (v1_waybel_0(k2_tarski(esk2_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_19])])).
cnf(c_0_65, negated_conjecture, (v2_waybel_0(k2_tarski(X1,X1),esk1_0)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_20])]), c_0_34])).
cnf(c_0_66, negated_conjecture, (~v2_waybel_0(k2_tarski(esk2_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_64])])).
cnf(c_0_67, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_60]), c_0_19])]), c_0_66]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 68
# Proof object clause steps            : 45
# Proof object formula steps           : 23
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 47
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 36
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 35
# Processed clauses                    : 110
# ...of these trivial                  : 0
# ...subsumed                          : 6
# ...remaining for further processing  : 104
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 4
# Generated clauses                    : 207
# ...of the previous two non-trivial   : 198
# Contextual simplify-reflections      : 2
# Paramodulations                      : 204
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 61
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 51
# Current number of unprocessed clauses: 154
# ...number of literals in the above   : 998
# Current number of archived formulas  : 0
# Current number of archived clauses   : 42
# Clause-clause subsumption calls (NU) : 1409
# Rec. Clause-clause subsumption calls : 213
# Non-unit clause-clause subsumptions  : 10
# Unit Clause-clause subsumption calls : 12
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8865

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.006 s
# Total time               : 0.015 s
# Maximum resident set size: 28672 pages
