# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/62e2c06e6027426e85f16970ae386438.E_input', d10_xboole_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/62e2c06e6027426e85f16970ae386438.E_input', t3_subset)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/62e2c06e6027426e85f16970ae386438.E_input', t4_subset)).
fof(t22_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)=>r1_tarski(k6_waybel_0(X1,X3),k6_waybel_0(X1,X2)))))), file('data/MPTP/proofs/62e2c06e6027426e85f16970ae386438.E_input', t22_waybel_0)).
fof(t21_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)=>r1_tarski(k5_waybel_0(X1,X2),k5_waybel_0(X1,X3)))))), file('data/MPTP/proofs/62e2c06e6027426e85f16970ae386438.E_input', t21_waybel_0)).
fof(t18_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k6_waybel_0(X1,X2))<=>r1_orders_2(X1,X2,X3))))), file('data/MPTP/proofs/62e2c06e6027426e85f16970ae386438.E_input', t18_waybel_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/62e2c06e6027426e85f16970ae386438.E_input', d3_tarski)).
fof(t17_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2))))), file('data/MPTP/proofs/62e2c06e6027426e85f16970ae386438.E_input', t17_waybel_0)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/62e2c06e6027426e85f16970ae386438.E_input', t5_subset)).
fof(dt_k6_waybel_0, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k6_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/62e2c06e6027426e85f16970ae386438.E_input', dt_k6_waybel_0)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/62e2c06e6027426e85f16970ae386438.E_input', t2_subset)).
fof(c_0_11, plain, ![X4, X5]:(((r1_tarski(X4,X5)|X4!=X5)&(r1_tarski(X5,X4)|X4!=X5))&(~r1_tarski(X4,X5)|~r1_tarski(X5,X4)|X4=X5)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_12, plain, ![X6, X7]:((~m1_subset_1(X6,k1_zfmisc_1(X7))|r1_tarski(X6,X7))&(~r1_tarski(X6,X7)|m1_subset_1(X6,k1_zfmisc_1(X7)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_13, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_14, plain, ![X8, X9, X10]:(~r2_hidden(X8,X9)|~m1_subset_1(X9,k1_zfmisc_1(X10))|m1_subset_1(X8,X10)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_15, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_13])).
fof(c_0_17, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)=>r1_tarski(k6_waybel_0(X1,X3),k6_waybel_0(X1,X2))))))), inference(assume_negation,[status(cth)],[t22_waybel_0])).
fof(c_0_18, plain, ![X26, X27, X28]:(v2_struct_0(X26)|(~v4_orders_2(X26)|~l1_orders_2(X26))|(~m1_subset_1(X27,u1_struct_0(X26))|(~m1_subset_1(X28,u1_struct_0(X26))|(~r1_orders_2(X26,X27,X28)|r1_tarski(k5_waybel_0(X26,X27),k5_waybel_0(X26,X28)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t21_waybel_0])])])])).
cnf(c_0_19, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_15, c_0_16])).
fof(c_0_21, plain, ![X20, X21, X22]:((~r2_hidden(X22,k6_waybel_0(X20,X21))|r1_orders_2(X20,X21,X22)|~m1_subset_1(X22,u1_struct_0(X20))|~m1_subset_1(X21,u1_struct_0(X20))|(v2_struct_0(X20)|~l1_orders_2(X20)))&(~r1_orders_2(X20,X21,X22)|r2_hidden(X22,k6_waybel_0(X20,X21))|~m1_subset_1(X22,u1_struct_0(X20))|~m1_subset_1(X21,u1_struct_0(X20))|(v2_struct_0(X20)|~l1_orders_2(X20)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t18_waybel_0])])])])])).
fof(c_0_22, negated_conjecture, ((~v2_struct_0(esk2_0)&(v4_orders_2(esk2_0)&l1_orders_2(esk2_0)))&(m1_subset_1(esk3_0,u1_struct_0(esk2_0))&(m1_subset_1(esk4_0,u1_struct_0(esk2_0))&(r1_orders_2(esk2_0,esk3_0,esk4_0)&~r1_tarski(k6_waybel_0(esk2_0,esk4_0),k6_waybel_0(esk2_0,esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])).
fof(c_0_23, plain, ![X11, X12, X13, X14, X15]:((~r1_tarski(X11,X12)|(~r2_hidden(X13,X11)|r2_hidden(X13,X12)))&((r2_hidden(esk1_2(X14,X15),X14)|r1_tarski(X14,X15))&(~r2_hidden(esk1_2(X14,X15),X15)|r1_tarski(X14,X15)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_24, plain, (v2_struct_0(X1)|r1_tarski(k5_waybel_0(X1,X2),k5_waybel_0(X1,X3))|~v4_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
fof(c_0_26, plain, ![X23, X24, X25]:((~r2_hidden(X25,k5_waybel_0(X23,X24))|r1_orders_2(X23,X25,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~m1_subset_1(X24,u1_struct_0(X23))|(v2_struct_0(X23)|~l1_orders_2(X23)))&(~r1_orders_2(X23,X25,X24)|r2_hidden(X25,k5_waybel_0(X23,X24))|~m1_subset_1(X25,u1_struct_0(X23))|~m1_subset_1(X24,u1_struct_0(X23))|(v2_struct_0(X23)|~l1_orders_2(X23)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_waybel_0])])])])])).
fof(c_0_27, plain, ![X31, X32, X33]:(~r2_hidden(X31,X32)|(~m1_subset_1(X32,k1_zfmisc_1(X33))|~v1_xboole_0(X33))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_28, plain, ![X29, X30]:(v2_struct_0(X29)|~l1_orders_2(X29)|~m1_subset_1(X30,u1_struct_0(X29))|m1_subset_1(k6_waybel_0(X29,X30),k1_zfmisc_1(u1_struct_0(X29)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_waybel_0])])])).
cnf(c_0_29, plain, (r2_hidden(X3,k6_waybel_0(X1,X2))|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, negated_conjecture, (l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_34, plain, (r2_hidden(esk1_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_35, plain, (v2_struct_0(X1)|r1_tarski(k5_waybel_0(X1,X2),k5_waybel_0(X1,X3))|~r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~v4_orders_2(X1)|~r2_hidden(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_36, negated_conjecture, (v4_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_37, plain, (r2_hidden(X2,k5_waybel_0(X1,X3))|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_38, plain, ![X34, X35]:(~m1_subset_1(X34,X35)|(v1_xboole_0(X35)|r2_hidden(X34,X35))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_39, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_40, plain, (v2_struct_0(X1)|m1_subset_1(k6_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_41, negated_conjecture, (r2_hidden(esk4_0,k6_waybel_0(esk2_0,X1))|~r1_orders_2(esk2_0,X1,esk4_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])]), c_0_32])).
cnf(c_0_42, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_43, negated_conjecture, (r1_orders_2(esk2_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_44, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_45, plain, (r2_hidden(esk1_2(X1,X2),X1)|r2_hidden(X3,X2)|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_46, plain, (r1_orders_2(X2,X1,X3)|v2_struct_0(X2)|~r2_hidden(X1,k5_waybel_0(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_47, negated_conjecture, (r1_tarski(k5_waybel_0(esk2_0,esk4_0),k5_waybel_0(esk2_0,X1))|~r1_orders_2(esk2_0,esk4_0,X1)|~r2_hidden(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_30]), c_0_31]), c_0_36])]), c_0_32])).
cnf(c_0_48, negated_conjecture, (r2_hidden(X1,k5_waybel_0(esk2_0,esk4_0))|~r1_orders_2(esk2_0,X1,esk4_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_30]), c_0_31])]), c_0_32])).
cnf(c_0_49, plain, (r1_orders_2(X2,X3,X1)|v2_struct_0(X2)|~r2_hidden(X1,k6_waybel_0(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_50, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_51, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(X2,k6_waybel_0(X1,X3))|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk4_0,k6_waybel_0(esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43])])).
cnf(c_0_53, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_33, c_0_44])).
cnf(c_0_54, negated_conjecture, (~r1_tarski(k6_waybel_0(esk2_0,esk4_0),k6_waybel_0(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_55, plain, (r2_hidden(esk1_2(X1,X2),X3)|r2_hidden(esk1_2(X1,X3),X1)|r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_45, c_0_34])).
cnf(c_0_56, plain, (v2_struct_0(X1)|r2_hidden(X2,k6_waybel_0(X1,X3))|~r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~r2_hidden(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_29, c_0_25])).
cnf(c_0_57, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~l1_orders_2(X1)|~r2_hidden(X2,k5_waybel_0(X1,X3))|~r2_hidden(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_46, c_0_25])).
cnf(c_0_58, negated_conjecture, (r2_hidden(X1,k5_waybel_0(esk2_0,X2))|~r1_orders_2(esk2_0,esk4_0,X2)|~r2_hidden(X1,k5_waybel_0(esk2_0,esk4_0))|~r2_hidden(X2,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_33, c_0_47])).
cnf(c_0_59, negated_conjecture, (r2_hidden(esk3_0,k5_waybel_0(esk2_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_42]), c_0_43])])).
cnf(c_0_60, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~l1_orders_2(X1)|~r2_hidden(X3,k6_waybel_0(X1,X2))|~r2_hidden(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_49, c_0_25])).
cnf(c_0_61, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|r2_hidden(esk4_0,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_50, c_0_30])).
cnf(c_0_62, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_31]), c_0_42])]), c_0_32])).
cnf(c_0_63, plain, (v2_struct_0(X1)|m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(X2,k6_waybel_0(X1,X3))|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_19, c_0_40])).
cnf(c_0_64, plain, (v2_struct_0(X1)|r2_hidden(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(X2,k6_waybel_0(X1,X3))|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_53, c_0_40])).
cnf(c_0_65, negated_conjecture, (r2_hidden(esk1_2(k6_waybel_0(esk2_0,esk4_0),X1),k6_waybel_0(esk2_0,esk4_0))|r2_hidden(esk1_2(k6_waybel_0(esk2_0,esk4_0),k6_waybel_0(esk2_0,esk3_0)),X1)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_66, negated_conjecture, (r2_hidden(X1,k6_waybel_0(esk2_0,esk3_0))|~r1_orders_2(esk2_0,esk3_0,X1)|~r2_hidden(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_42]), c_0_31])]), c_0_32])).
cnf(c_0_67, negated_conjecture, (r1_orders_2(esk2_0,esk3_0,X1)|~r2_hidden(esk3_0,k5_waybel_0(esk2_0,X1))|~r2_hidden(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_42]), c_0_31])]), c_0_32])).
cnf(c_0_68, negated_conjecture, (r2_hidden(esk3_0,k5_waybel_0(esk2_0,X1))|~r1_orders_2(esk2_0,esk4_0,X1)|~r2_hidden(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_69, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~l1_orders_2(X1)|~r2_hidden(X3,k6_waybel_0(X1,X2))|~r2_hidden(X2,u1_struct_0(X1))|~r2_hidden(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_60, c_0_25])).
cnf(c_0_70, negated_conjecture, (r2_hidden(esk4_0,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_71, plain, (v2_struct_0(X1)|r2_hidden(esk1_2(X2,k6_waybel_0(X1,X3)),X2)|m1_subset_1(esk1_2(X2,X4),u1_struct_0(X1))|r1_tarski(X2,X4)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_63, c_0_55])).
cnf(c_0_72, plain, (r1_tarski(X1,X2)|~r2_hidden(esk1_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_73, negated_conjecture, (r2_hidden(esk1_2(k6_waybel_0(esk2_0,esk4_0),k6_waybel_0(esk2_0,esk3_0)),X1)|r2_hidden(esk1_2(k6_waybel_0(esk2_0,esk4_0),X1),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_31]), c_0_30])]), c_0_32])).
cnf(c_0_74, negated_conjecture, (r2_hidden(X1,k6_waybel_0(esk2_0,esk3_0))|~r2_hidden(esk3_0,k5_waybel_0(esk2_0,X1))|~r2_hidden(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_75, negated_conjecture, (r2_hidden(esk3_0,k5_waybel_0(esk2_0,X1))|~r2_hidden(X1,k6_waybel_0(esk2_0,esk4_0))|~r2_hidden(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_31]), c_0_70])]), c_0_32])).
cnf(c_0_76, negated_conjecture, (r2_hidden(esk1_2(X1,k6_waybel_0(esk2_0,esk3_0)),X1)|m1_subset_1(esk1_2(X1,X2),u1_struct_0(esk2_0))|r1_tarski(X1,X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_42]), c_0_31])]), c_0_32])).
cnf(c_0_77, negated_conjecture, (r2_hidden(esk1_2(k6_waybel_0(esk2_0,esk4_0),k6_waybel_0(esk2_0,esk3_0)),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_54])).
cnf(c_0_78, negated_conjecture, (r1_tarski(X1,k6_waybel_0(esk2_0,esk3_0))|~r2_hidden(esk3_0,k5_waybel_0(esk2_0,esk1_2(X1,k6_waybel_0(esk2_0,esk3_0))))|~r2_hidden(esk1_2(X1,k6_waybel_0(esk2_0,esk3_0)),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_72, c_0_74])).
cnf(c_0_79, negated_conjecture, (r2_hidden(esk3_0,k5_waybel_0(esk2_0,esk1_2(k6_waybel_0(esk2_0,esk4_0),k6_waybel_0(esk2_0,esk3_0))))|m1_subset_1(esk1_2(k6_waybel_0(esk2_0,esk4_0),X1),u1_struct_0(esk2_0))|r1_tarski(k6_waybel_0(esk2_0,esk4_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77])])).
cnf(c_0_80, negated_conjecture, (~r2_hidden(esk3_0,k5_waybel_0(esk2_0,esk1_2(k6_waybel_0(esk2_0,esk4_0),k6_waybel_0(esk2_0,esk3_0))))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_77]), c_0_54])).
cnf(c_0_81, negated_conjecture, (m1_subset_1(esk1_2(k6_waybel_0(esk2_0,esk4_0),X1),u1_struct_0(esk2_0))|r1_tarski(k6_waybel_0(esk2_0,esk4_0),X1)), inference(sr,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_82, negated_conjecture, (r2_hidden(esk1_2(k6_waybel_0(esk2_0,esk4_0),X1),u1_struct_0(esk2_0))|r1_tarski(k6_waybel_0(esk2_0,esk4_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_81]), c_0_62])).
cnf(c_0_83, negated_conjecture, (r1_tarski(k6_waybel_0(esk2_0,esk4_0),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_72, c_0_82])).
cnf(c_0_84, negated_conjecture, (m1_subset_1(k6_waybel_0(esk2_0,esk4_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_15, c_0_83])).
cnf(c_0_85, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,k6_waybel_0(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_19, c_0_84])).
cnf(c_0_86, negated_conjecture, (r1_orders_2(esk2_0,esk4_0,X1)|~r2_hidden(X1,k6_waybel_0(esk2_0,esk4_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_30]), c_0_31])]), c_0_32])).
cnf(c_0_87, negated_conjecture, (r1_tarski(k5_waybel_0(esk2_0,X1),k5_waybel_0(esk2_0,X2))|~r1_orders_2(esk2_0,X1,X2)|~r2_hidden(X2,k6_waybel_0(esk2_0,esk4_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_85]), c_0_31]), c_0_36])]), c_0_32])).
cnf(c_0_88, negated_conjecture, (r1_orders_2(esk2_0,esk4_0,X1)|~r2_hidden(X1,k6_waybel_0(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_86, c_0_85])).
cnf(c_0_89, negated_conjecture, (r1_tarski(k5_waybel_0(esk2_0,esk4_0),k5_waybel_0(esk2_0,X1))|~r2_hidden(X1,k6_waybel_0(esk2_0,esk4_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_30]), c_0_88])).
cnf(c_0_90, negated_conjecture, (r2_hidden(X1,k5_waybel_0(esk2_0,X2))|~r2_hidden(X1,k5_waybel_0(esk2_0,esk4_0))|~r2_hidden(X2,k6_waybel_0(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_33, c_0_89])).
cnf(c_0_91, negated_conjecture, (r2_hidden(esk3_0,k5_waybel_0(esk2_0,X1))|~r2_hidden(X1,k6_waybel_0(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_90, c_0_59])).
cnf(c_0_92, negated_conjecture, (r2_hidden(esk3_0,k5_waybel_0(esk2_0,esk1_2(k6_waybel_0(esk2_0,esk4_0),X1)))|r1_tarski(k6_waybel_0(esk2_0,esk4_0),X1)), inference(spm,[status(thm)],[c_0_91, c_0_34])).
cnf(c_0_93, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_92]), c_0_80]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 94
# Proof object clause steps            : 71
# Proof object formula steps           : 23
# Proof object conjectures             : 44
# Proof object clause conjectures      : 41
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 46
# Proof object simplifying inferences  : 52
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 24
# Processed clauses                    : 3777
# ...of these trivial                  : 2
# ...subsumed                          : 2433
# ...remaining for further processing  : 1342
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 128
# Backward-rewritten                   : 9
# Generated clauses                    : 49892
# ...of the previous two non-trivial   : 47905
# Contextual simplify-reflections      : 13
# Paramodulations                      : 49775
# Factorizations                       : 102
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1173
#    Positive orientable unit clauses  : 22
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 12
#    Non-unit-clauses                  : 1139
# Current number of unprocessed clauses: 43988
# ...number of literals in the above   : 242188
# Current number of archived formulas  : 0
# Current number of archived clauses   : 164
# Clause-clause subsumption calls (NU) : 160806
# Rec. Clause-clause subsumption calls : 65942
# Non-unit clause-clause subsumptions  : 2195
# Unit Clause-clause subsumption calls : 2124
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 28
# BW rewrite match successes           : 8
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1047167

# -------------------------------------------------
# User time                : 0.656 s
# System time              : 0.013 s
# Total time               : 0.669 s
# Maximum resident set size: 35952 pages
