# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C41_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/6379ef7f462e4d1b8d00db4c1bfea764.E_input', d2_zfmisc_1)).
fof(t67_mcart_1, conjecture, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1)))), file('data/MPTP/proofs/6379ef7f462e4d1b8d00db4c1bfea764.E_input', t67_mcart_1)).
fof(t20_mcart_1, axiom, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1)))), file('data/MPTP/proofs/6379ef7f462e4d1b8d00db4c1bfea764.E_input', t20_mcart_1)).
fof(c_0_3, plain, ![X13, X14, X15, X16, X19, X20, X21, X22, X23, X24, X26, X27]:((((r2_hidden(esk4_4(X13,X14,X15,X16),X13)|~r2_hidden(X16,X15)|X15!=k2_zfmisc_1(X13,X14))&((r2_hidden(esk5_4(X13,X14,X15,X16),X14)|~r2_hidden(X16,X15)|X15!=k2_zfmisc_1(X13,X14))&(X16=k4_tarski(esk4_4(X13,X14,X15,X16),esk5_4(X13,X14,X15,X16))|~r2_hidden(X16,X15)|X15!=k2_zfmisc_1(X13,X14))))&(~r2_hidden(X20,X13)|(~r2_hidden(X21,X14)|X19!=k4_tarski(X20,X21))|r2_hidden(X19,X15)|X15!=k2_zfmisc_1(X13,X14)))&((~r2_hidden(esk6_3(X22,X23,X24),X24)|(~r2_hidden(X26,X22)|(~r2_hidden(X27,X23)|esk6_3(X22,X23,X24)!=k4_tarski(X26,X27)))|X24=k2_zfmisc_1(X22,X23))&((r2_hidden(esk7_3(X22,X23,X24),X22)|r2_hidden(esk6_3(X22,X23,X24),X24)|X24=k2_zfmisc_1(X22,X23))&((r2_hidden(esk8_3(X22,X23,X24),X23)|r2_hidden(esk6_3(X22,X23,X24),X24)|X24=k2_zfmisc_1(X22,X23))&(esk6_3(X22,X23,X24)=k4_tarski(esk7_3(X22,X23,X24),esk8_3(X22,X23,X24))|r2_hidden(esk6_3(X22,X23,X24),X24)|X24=k2_zfmisc_1(X22,X23)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
fof(c_0_4, negated_conjecture, ~(![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1))))), inference(assume_negation,[status(cth)],[t67_mcart_1])).
fof(c_0_5, plain, ![X10, X11, X12]:((X10!=k1_mcart_1(X10)|X10!=k4_tarski(X11,X12))&(X10!=k2_mcart_1(X10)|X10!=k4_tarski(X11,X12))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t20_mcart_1])])])])])).
cnf(c_0_6, plain, (X1=k4_tarski(esk4_4(X2,X3,X4,X1),esk5_4(X2,X3,X4,X1))|~r2_hidden(X1,X4)|X4!=k2_zfmisc_1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_3])).
fof(c_0_7, negated_conjecture, (r2_hidden(esk1_0,k2_zfmisc_1(esk2_0,esk3_0))&(esk1_0=k1_mcart_1(esk1_0)|esk1_0=k2_mcart_1(esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_4])])])])).
cnf(c_0_8, plain, (X1!=k1_mcart_1(X1)|X1!=k4_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (k4_tarski(esk4_4(X1,X2,k2_zfmisc_1(X1,X2),X3),esk5_4(X1,X2,k2_zfmisc_1(X1,X2),X3))=X3|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_6])).
cnf(c_0_10, negated_conjecture, (r2_hidden(esk1_0,k2_zfmisc_1(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (X1!=k2_mcart_1(X1)|X1!=k4_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_12, plain, (k1_mcart_1(k4_tarski(X1,X2))!=k4_tarski(X1,X2)), inference(er,[status(thm)],[c_0_8])).
cnf(c_0_13, negated_conjecture, (k4_tarski(esk4_4(esk2_0,esk3_0,k2_zfmisc_1(esk2_0,esk3_0),esk1_0),esk5_4(esk2_0,esk3_0,k2_zfmisc_1(esk2_0,esk3_0),esk1_0))=esk1_0), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_14, plain, (k2_mcart_1(k4_tarski(X1,X2))!=k4_tarski(X1,X2)), inference(er,[status(thm)],[c_0_11])).
cnf(c_0_15, negated_conjecture, (esk1_0=k1_mcart_1(esk1_0)|esk1_0=k2_mcart_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_16, negated_conjecture, (k1_mcart_1(esk1_0)!=esk1_0), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_17, negated_conjecture, (k2_mcart_1(esk1_0)!=esk1_0), inference(spm,[status(thm)],[c_0_14, c_0_13])).
cnf(c_0_18, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_15, c_0_16]), c_0_17]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 19
# Proof object clause steps            : 12
# Proof object formula steps           : 7
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 5
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 2
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 12
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 12
# Processed clauses                    : 45
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 41
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 137
# ...of the previous two non-trivial   : 137
# Contextual simplify-reflections      : 0
# Paramodulations                      : 126
# Factorizations                       : 0
# Equation resolutions                 : 10
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 40
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 24
# Current number of unprocessed clauses: 104
# ...number of literals in the above   : 198
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1
# Clause-clause subsumption calls (NU) : 65
# Rec. Clause-clause subsumption calls : 34
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 46
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4005

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 26144 pages
