# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t149_zfmisc_1, conjecture, ![X1, X2, X3, X4]:((r1_tarski(k3_xboole_0(X1,X2),k1_tarski(X4))&(r2_hidden(X4,X3)&r1_xboole_0(X3,X2)))=>r1_xboole_0(k2_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/63c26286b83d4cdb898570b76288822a.E_input', t149_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/63c26286b83d4cdb898570b76288822a.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/63c26286b83d4cdb898570b76288822a.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/63c26286b83d4cdb898570b76288822a.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/63c26286b83d4cdb898570b76288822a.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/63c26286b83d4cdb898570b76288822a.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/63c26286b83d4cdb898570b76288822a.E_input', t74_enumset1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/63c26286b83d4cdb898570b76288822a.E_input', t28_xboole_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/63c26286b83d4cdb898570b76288822a.E_input', t3_xboole_0)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/63c26286b83d4cdb898570b76288822a.E_input', l51_zfmisc_1)).
fof(t16_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3)), file('data/MPTP/proofs/63c26286b83d4cdb898570b76288822a.E_input', t16_xboole_1)).
fof(t66_xboole_1, axiom, ![X1]:(~((~(r1_xboole_0(X1,X1))&X1=k1_xboole_0))&~((~(X1=k1_xboole_0)&r1_xboole_0(X1,X1)))), file('data/MPTP/proofs/63c26286b83d4cdb898570b76288822a.E_input', t66_xboole_1)).
fof(t76_xboole_1, axiom, ![X1, X2, X3]:(r1_xboole_0(X1,X2)=>r1_xboole_0(k3_xboole_0(X3,X1),k3_xboole_0(X3,X2))), file('data/MPTP/proofs/63c26286b83d4cdb898570b76288822a.E_input', t76_xboole_1)).
fof(l27_zfmisc_1, axiom, ![X1, X2]:(~(r2_hidden(X1,X2))=>r1_xboole_0(k1_tarski(X1),X2)), file('data/MPTP/proofs/63c26286b83d4cdb898570b76288822a.E_input', l27_zfmisc_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/63c26286b83d4cdb898570b76288822a.E_input', d7_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/63c26286b83d4cdb898570b76288822a.E_input', commutativity_k3_xboole_0)).
fof(t78_xboole_1, axiom, ![X1, X2, X3]:(r1_xboole_0(X1,X2)=>k3_xboole_0(X1,k2_xboole_0(X2,X3))=k3_xboole_0(X1,X3)), file('data/MPTP/proofs/63c26286b83d4cdb898570b76288822a.E_input', t78_xboole_1)).
fof(c_0_17, negated_conjecture, ~(![X1, X2, X3, X4]:((r1_tarski(k3_xboole_0(X1,X2),k1_tarski(X4))&(r2_hidden(X4,X3)&r1_xboole_0(X3,X2)))=>r1_xboole_0(k2_xboole_0(X1,X3),X2))), inference(assume_negation,[status(cth)],[t149_zfmisc_1])).
fof(c_0_18, negated_conjecture, ((r1_tarski(k3_xboole_0(esk1_0,esk2_0),k1_tarski(esk4_0))&(r2_hidden(esk4_0,esk3_0)&r1_xboole_0(esk3_0,esk2_0)))&~r1_xboole_0(k2_xboole_0(esk1_0,esk3_0),esk2_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).
fof(c_0_19, plain, ![X18]:k2_tarski(X18,X18)=k1_tarski(X18), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_20, plain, ![X19, X20]:k1_enumset1(X19,X19,X20)=k2_tarski(X19,X20), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_21, plain, ![X21, X22, X23]:k2_enumset1(X21,X21,X22,X23)=k1_enumset1(X21,X22,X23), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_22, plain, ![X24, X25, X26, X27]:k3_enumset1(X24,X24,X25,X26,X27)=k2_enumset1(X24,X25,X26,X27), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_23, plain, ![X28, X29, X30, X31, X32]:k4_enumset1(X28,X28,X29,X30,X31,X32)=k3_enumset1(X28,X29,X30,X31,X32), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_24, plain, ![X33, X34, X35, X36, X37, X38]:k5_enumset1(X33,X33,X34,X35,X36,X37,X38)=k4_enumset1(X33,X34,X35,X36,X37,X38), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_25, plain, ![X42, X43]:(~r1_tarski(X42,X43)|k3_xboole_0(X42,X43)=X42), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_26, negated_conjecture, (r1_tarski(k3_xboole_0(esk1_0,esk2_0),k1_tarski(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_33, plain, ![X12, X13, X15, X16, X17]:(((r2_hidden(esk5_2(X12,X13),X12)|r1_xboole_0(X12,X13))&(r2_hidden(esk5_2(X12,X13),X13)|r1_xboole_0(X12,X13)))&(~r2_hidden(X17,X15)|~r2_hidden(X17,X16)|~r1_xboole_0(X15,X16))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_34, plain, ![X53, X54]:k3_tarski(k2_tarski(X53,X54))=k2_xboole_0(X53,X54), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
cnf(c_0_35, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_36, negated_conjecture, (r1_tarski(k3_xboole_0(esk1_0,esk2_0),k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32])).
fof(c_0_37, plain, ![X44, X45, X46]:k3_xboole_0(k3_xboole_0(X44,X45),X46)=k3_xboole_0(X44,k3_xboole_0(X45,X46)), inference(variable_rename,[status(thm)],[t16_xboole_1])).
cnf(c_0_38, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_39, plain, (r2_hidden(esk5_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_40, plain, ![X7]:((r1_xboole_0(X7,X7)|X7!=k1_xboole_0)&(X7=k1_xboole_0|~r1_xboole_0(X7,X7))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t66_xboole_1])])])).
cnf(c_0_41, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_42, plain, ![X39, X40, X41]:(~r1_xboole_0(X39,X40)|r1_xboole_0(k3_xboole_0(X41,X39),k3_xboole_0(X41,X40))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t76_xboole_1])])).
cnf(c_0_43, negated_conjecture, (k3_xboole_0(k3_xboole_0(esk1_0,esk2_0),k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0))=k3_xboole_0(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_44, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_45, plain, ![X51, X52]:(r2_hidden(X51,X52)|r1_xboole_0(k1_tarski(X51),X52)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l27_zfmisc_1])])])).
cnf(c_0_46, plain, (r1_xboole_0(X1,X2)|~r2_hidden(esk5_2(X1,X2),X3)|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_47, plain, (r1_xboole_0(X1,X1)|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_48, negated_conjecture, (~r1_xboole_0(k2_xboole_0(esk1_0,esk3_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_49, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_41, c_0_28])).
fof(c_0_50, plain, ![X49, X50]:((~r1_xboole_0(X49,X50)|k3_xboole_0(X49,X50)=k1_xboole_0)&(k3_xboole_0(X49,X50)!=k1_xboole_0|r1_xboole_0(X49,X50))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_51, plain, (r1_xboole_0(k3_xboole_0(X3,X1),k3_xboole_0(X3,X2))|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_52, negated_conjecture, (k3_xboole_0(esk1_0,k3_xboole_0(esk2_0,k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)))=k3_xboole_0(esk1_0,esk2_0)), inference(rw,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_53, plain, (r2_hidden(X1,X2)|r1_xboole_0(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_54, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_46, c_0_39])).
cnf(c_0_55, plain, (r1_xboole_0(k1_xboole_0,k1_xboole_0)), inference(er,[status(thm)],[c_0_47])).
cnf(c_0_56, negated_conjecture, (~r1_xboole_0(k3_tarski(k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0)),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_49]), c_0_29]), c_0_30]), c_0_31]), c_0_32])).
cnf(c_0_57, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_50])).
fof(c_0_58, plain, ![X47, X48]:k3_xboole_0(X47,X48)=k3_xboole_0(X48,X47), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_59, plain, ![X55, X56, X57]:(~r1_xboole_0(X55,X56)|k3_xboole_0(X55,k2_xboole_0(X56,X57))=k3_xboole_0(X55,X57)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t78_xboole_1])])).
cnf(c_0_60, negated_conjecture, (r2_hidden(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_61, negated_conjecture, (r1_xboole_0(k3_xboole_0(esk1_0,X1),k3_xboole_0(esk1_0,esk2_0))|~r1_xboole_0(X1,k3_xboole_0(esk2_0,k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)))), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_62, plain, (r2_hidden(X1,X2)|r1_xboole_0(k5_enumset1(X1,X1,X1,X1,X1,X1,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32])).
cnf(c_0_63, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_64, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_65, negated_conjecture, (r1_xboole_0(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_66, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_67, negated_conjecture, (k3_xboole_0(k3_tarski(k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0)),esk2_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_68, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_69, plain, (k3_xboole_0(X1,k2_xboole_0(X2,X3))=k3_xboole_0(X1,X3)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_70, negated_conjecture, (~r2_hidden(esk4_0,X1)|~r1_xboole_0(X1,esk3_0)), inference(spm,[status(thm)],[c_0_38, c_0_60])).
cnf(c_0_71, negated_conjecture, (r2_hidden(X1,k3_xboole_0(esk2_0,k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)))|r1_xboole_0(k3_xboole_0(esk1_0,k5_enumset1(X1,X1,X1,X1,X1,X1,X1)),k3_xboole_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_72, plain, (r1_xboole_0(X1,X2)|~r2_hidden(esk5_2(X1,X2),X3)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_38, c_0_63])).
cnf(c_0_73, negated_conjecture, (k3_xboole_0(esk3_0,esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_74, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_64, c_0_66])).
cnf(c_0_75, plain, (X1=k1_xboole_0|~r1_xboole_0(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_76, negated_conjecture, (k3_xboole_0(esk2_0,k3_tarski(k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0)))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_77, plain, (k3_xboole_0(X1,k3_tarski(k5_enumset1(X2,X2,X2,X2,X2,X2,X3)))=k3_xboole_0(X1,X3)|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_49]), c_0_29]), c_0_30]), c_0_31]), c_0_32])).
cnf(c_0_78, negated_conjecture, (r1_xboole_0(k3_xboole_0(esk1_0,k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)),k3_xboole_0(esk1_0,esk2_0))|~r1_xboole_0(k3_xboole_0(esk2_0,k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)),esk3_0)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_79, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(X2,k3_xboole_0(X3,X1))), inference(spm,[status(thm)],[c_0_44, c_0_68])).
cnf(c_0_80, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_72, c_0_39])).
cnf(c_0_81, negated_conjecture, (k3_xboole_0(esk3_0,k3_xboole_0(esk2_0,X1))=k3_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_44, c_0_73])).
cnf(c_0_82, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_68, c_0_74])).
cnf(c_0_83, plain, (X1=k1_xboole_0|k3_xboole_0(X1,X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_75, c_0_57])).
cnf(c_0_84, negated_conjecture, (~r1_xboole_0(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_68]), c_0_73])])).
cnf(c_0_85, negated_conjecture, (k3_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0))=k1_xboole_0|~r1_xboole_0(k3_xboole_0(esk2_0,k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)),esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_78]), c_0_44]), c_0_79]), c_0_52])).
cnf(c_0_86, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X2,X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_80, c_0_57])).
cnf(c_0_87, negated_conjecture, (k3_xboole_0(esk3_0,k3_xboole_0(esk2_0,X1))=k1_xboole_0), inference(rw,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_88, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|k3_xboole_0(X1,k3_xboole_0(X2,k3_xboole_0(X1,X2)))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_83, c_0_44])).
cnf(c_0_89, negated_conjecture, (k3_xboole_0(esk1_0,k3_xboole_0(esk2_0,k3_xboole_0(k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0),X1)))=k3_xboole_0(esk1_0,k3_xboole_0(esk2_0,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_52]), c_0_44]), c_0_44])).
cnf(c_0_90, negated_conjecture, (k3_xboole_0(esk1_0,esk2_0)!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_57]), c_0_68])).
cnf(c_0_91, negated_conjecture, (k3_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_87])])).
cnf(c_0_92, negated_conjecture, (k3_xboole_0(esk1_0,k3_xboole_0(esk1_0,k3_xboole_0(esk2_0,esk2_0)))!=k1_xboole_0), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_52]), c_0_44]), c_0_89]), c_0_90]), c_0_79])).
cnf(c_0_93, negated_conjecture, (k3_xboole_0(esk1_0,k3_xboole_0(esk1_0,k3_xboole_0(esk2_0,X1)))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_91]), c_0_82]), c_0_44])).
cnf(c_0_94, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_93])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 95
# Proof object clause steps            : 60
# Proof object formula steps           : 35
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 26
# Proof object simplifying inferences  : 46
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 17
# Processed clauses                    : 364
# ...of these trivial                  : 53
# ...subsumed                          : 154
# ...remaining for further processing  : 157
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 9
# Generated clauses                    : 1743
# ...of the previous two non-trivial   : 1027
# Contextual simplify-reflections      : 0
# Paramodulations                      : 1742
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 129
#    Positive orientable unit clauses  : 29
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 91
# Current number of unprocessed clauses: 687
# ...number of literals in the above   : 1283
# Current number of archived formulas  : 0
# Current number of archived clauses   : 34
# Clause-clause subsumption calls (NU) : 2761
# Rec. Clause-clause subsumption calls : 2756
# Non-unit clause-clause subsumptions  : 139
# Unit Clause-clause subsumption calls : 69
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 102
# BW rewrite match successes           : 78
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 19992

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.009 s
# Total time               : 0.020 s
# Maximum resident set size: 28648 pages
