# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t39_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2))), file('data/MPTP/proofs/64c57700c20f404e98bd4ef224524f8e.E_input', t39_yellow_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/64c57700c20f404e98bd4ef224524f8e.E_input', t24_orders_2)).
fof(t19_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))))), file('data/MPTP/proofs/64c57700c20f404e98bd4ef224524f8e.E_input', t19_yellow_0)).
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/64c57700c20f404e98bd4ef224524f8e.E_input', d9_yellow_0)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/64c57700c20f404e98bd4ef224524f8e.E_input', dt_k1_yellow_0)).
fof(t7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2)))))))), file('data/MPTP/proofs/64c57700c20f404e98bd4ef224524f8e.E_input', t7_yellow_0)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/64c57700c20f404e98bd4ef224524f8e.E_input', t69_enumset1)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/64c57700c20f404e98bd4ef224524f8e.E_input', redefinition_k6_domain_1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/64c57700c20f404e98bd4ef224524f8e.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/64c57700c20f404e98bd4ef224524f8e.E_input', dt_l1_orders_2)).
fof(t38_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))&r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))))), file('data/MPTP/proofs/64c57700c20f404e98bd4ef224524f8e.E_input', t38_yellow_0)).
fof(d8_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X5)=>r1_orders_2(X1,X5,X4))))=>X4=X3))))))), file('data/MPTP/proofs/64c57700c20f404e98bd4ef224524f8e.E_input', d8_yellow_0)).
fof(t31_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3))=>(r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2)))))&((r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2))))=>(X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3)))))), file('data/MPTP/proofs/64c57700c20f404e98bd4ef224524f8e.E_input', t31_yellow_0)).
fof(c_0_13, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))), introduced(definition)).
fof(c_0_14, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))), inference(split_equiv,[status(thm)],[c_0_13])).
fof(c_0_15, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2)))), inference(assume_negation,[status(cth)],[t39_yellow_0])).
fof(c_0_16, plain, ![X47, X48, X49, X50, X51]:(((r1_orders_2(X49,X50,X48)|(X50!=k11_lattice3(X49,X48,X47)|~r2_yellow_0(X49,k2_tarski(X48,X47)))|~m1_subset_1(X50,u1_struct_0(X49))|~epred1_3(X49,X48,X47))&((r1_orders_2(X49,X50,X47)|(X50!=k11_lattice3(X49,X48,X47)|~r2_yellow_0(X49,k2_tarski(X48,X47)))|~m1_subset_1(X50,u1_struct_0(X49))|~epred1_3(X49,X48,X47))&(~m1_subset_1(X51,u1_struct_0(X49))|(~r1_orders_2(X49,X51,X48)|~r1_orders_2(X49,X51,X47)|r1_orders_2(X49,X51,X50))|(X50!=k11_lattice3(X49,X48,X47)|~r2_yellow_0(X49,k2_tarski(X48,X47)))|~m1_subset_1(X50,u1_struct_0(X49))|~epred1_3(X49,X48,X47))))&(((X50=k11_lattice3(X49,X48,X47)|(m1_subset_1(esk9_4(X47,X48,X49,X50),u1_struct_0(X49))|~r1_orders_2(X49,X50,X47)|~r1_orders_2(X49,X50,X48))|~m1_subset_1(X50,u1_struct_0(X49))|~epred1_3(X49,X48,X47))&(r2_yellow_0(X49,k2_tarski(X48,X47))|(m1_subset_1(esk9_4(X47,X48,X49,X50),u1_struct_0(X49))|~r1_orders_2(X49,X50,X47)|~r1_orders_2(X49,X50,X48))|~m1_subset_1(X50,u1_struct_0(X49))|~epred1_3(X49,X48,X47)))&((((X50=k11_lattice3(X49,X48,X47)|(r1_orders_2(X49,esk9_4(X47,X48,X49,X50),X48)|~r1_orders_2(X49,X50,X47)|~r1_orders_2(X49,X50,X48))|~m1_subset_1(X50,u1_struct_0(X49))|~epred1_3(X49,X48,X47))&(r2_yellow_0(X49,k2_tarski(X48,X47))|(r1_orders_2(X49,esk9_4(X47,X48,X49,X50),X48)|~r1_orders_2(X49,X50,X47)|~r1_orders_2(X49,X50,X48))|~m1_subset_1(X50,u1_struct_0(X49))|~epred1_3(X49,X48,X47)))&((X50=k11_lattice3(X49,X48,X47)|(r1_orders_2(X49,esk9_4(X47,X48,X49,X50),X47)|~r1_orders_2(X49,X50,X47)|~r1_orders_2(X49,X50,X48))|~m1_subset_1(X50,u1_struct_0(X49))|~epred1_3(X49,X48,X47))&(r2_yellow_0(X49,k2_tarski(X48,X47))|(r1_orders_2(X49,esk9_4(X47,X48,X49,X50),X47)|~r1_orders_2(X49,X50,X47)|~r1_orders_2(X49,X50,X48))|~m1_subset_1(X50,u1_struct_0(X49))|~epred1_3(X49,X48,X47))))&((X50=k11_lattice3(X49,X48,X47)|(~r1_orders_2(X49,esk9_4(X47,X48,X49,X50),X50)|~r1_orders_2(X49,X50,X47)|~r1_orders_2(X49,X50,X48))|~m1_subset_1(X50,u1_struct_0(X49))|~epred1_3(X49,X48,X47))&(r2_yellow_0(X49,k2_tarski(X48,X47))|(~r1_orders_2(X49,esk9_4(X47,X48,X49,X50),X50)|~r1_orders_2(X49,X50,X47)|~r1_orders_2(X49,X50,X48))|~m1_subset_1(X50,u1_struct_0(X49))|~epred1_3(X49,X48,X47)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])])])).
fof(c_0_17, plain, ![X41, X42]:(v2_struct_0(X41)|(~v3_orders_2(X41)|~l1_orders_2(X41))|(~m1_subset_1(X42,u1_struct_0(X41))|r1_orders_2(X41,X42,X42))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_18, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_15])])])])).
cnf(c_0_19, plain, (X1=k11_lattice3(X2,X3,X4)|~r1_orders_2(X2,esk9_4(X4,X3,X2,X1),X1)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, (X1=k11_lattice3(X2,X3,X4)|r1_orders_2(X2,esk9_4(X4,X3,X2,X1),X4)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (X1=k11_lattice3(X2,X3,X4)|r1_orders_2(X2,esk9_4(X4,X3,X2,X1),X3)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26, plain, (k11_lattice3(X1,X2,X3)=X3|~epred1_3(X1,X2,X3)|~r1_orders_2(X1,X3,X3)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_27, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])]), c_0_24])).
cnf(c_0_28, plain, (k11_lattice3(X1,X2,X3)=X2|~epred1_3(X1,X2,X3)|~r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_19, c_0_25])).
fof(c_0_29, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>epred1_3(X1,X2,X3)))), inference(apply_def,[status(thm)],[t19_yellow_0, c_0_13])).
cnf(c_0_30, negated_conjecture, (k11_lattice3(esk1_0,X1,X2)=X2|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_31, negated_conjecture, (k11_lattice3(esk1_0,X1,X2)=X1|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_28, c_0_27])).
fof(c_0_32, plain, ![X32, X33, X34]:(~v5_orders_2(X32)|~l1_orders_2(X32)|(~m1_subset_1(X33,u1_struct_0(X32))|(~m1_subset_1(X34,u1_struct_0(X32))|epred1_3(X32,X33,X34)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])])).
fof(c_0_33, plain, ![X23, X24, X25, X26]:(((r2_lattice3(X23,X24,X25)|X25!=k1_yellow_0(X23,X24)|~r1_yellow_0(X23,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))&(~m1_subset_1(X26,u1_struct_0(X23))|(~r2_lattice3(X23,X24,X26)|r1_orders_2(X23,X25,X26))|X25!=k1_yellow_0(X23,X24)|~r1_yellow_0(X23,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23)))&((m1_subset_1(esk7_3(X23,X24,X25),u1_struct_0(X23))|~r2_lattice3(X23,X24,X25)|X25=k1_yellow_0(X23,X24)|~r1_yellow_0(X23,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))&((r2_lattice3(X23,X24,esk7_3(X23,X24,X25))|~r2_lattice3(X23,X24,X25)|X25=k1_yellow_0(X23,X24)|~r1_yellow_0(X23,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))&(~r1_orders_2(X23,X25,esk7_3(X23,X24,X25))|~r2_lattice3(X23,X24,X25)|X25=k1_yellow_0(X23,X24)|~r1_yellow_0(X23,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
fof(c_0_34, plain, ![X21, X22]:(~l1_orders_2(X21)|m1_subset_1(k1_yellow_0(X21,X22),u1_struct_0(X21))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
fof(c_0_35, plain, ![X29, X30, X31]:((~r1_lattice3(X29,k1_tarski(X31),X30)|r1_orders_2(X29,X30,X31)|~m1_subset_1(X31,u1_struct_0(X29))|~m1_subset_1(X30,u1_struct_0(X29))|~l1_orders_2(X29))&((~r1_orders_2(X29,X30,X31)|r1_lattice3(X29,k1_tarski(X31),X30)|~m1_subset_1(X31,u1_struct_0(X29))|~m1_subset_1(X30,u1_struct_0(X29))|~l1_orders_2(X29))&((~r2_lattice3(X29,k1_tarski(X31),X30)|r1_orders_2(X29,X31,X30)|~m1_subset_1(X31,u1_struct_0(X29))|~m1_subset_1(X30,u1_struct_0(X29))|~l1_orders_2(X29))&(~r1_orders_2(X29,X31,X30)|r2_lattice3(X29,k1_tarski(X31),X30)|~m1_subset_1(X31,u1_struct_0(X29))|~m1_subset_1(X30,u1_struct_0(X29))|~l1_orders_2(X29))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])).
fof(c_0_36, plain, ![X20]:k2_tarski(X20,X20)=k1_tarski(X20), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_37, plain, ![X45, X46]:(v1_xboole_0(X45)|~m1_subset_1(X46,X45)|k6_domain_1(X45,X46)=k1_tarski(X46)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_38, negated_conjecture, (X1=X2|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,X1)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_39, plain, (epred1_3(X1,X2,X3)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_41, plain, (r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~r1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43, plain, (r1_orders_2(X1,X2,X3)|~r2_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_44, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_45, plain, (r2_lattice3(X1,X2,X3)|X3!=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_46, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_47, negated_conjecture, (X1=X2|~r1_orders_2(esk1_0,X2,X1)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_23]), c_0_40])])).
cnf(c_0_48, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),X3)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_41]), c_0_42])).
cnf(c_0_49, plain, (r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r2_lattice3(X1,k2_tarski(X2,X2),X3)), inference(rw,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_50, plain, (r2_lattice3(X1,X2,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_45]), c_0_42])).
fof(c_0_51, plain, ![X28]:(v2_struct_0(X28)|~l1_struct_0(X28)|~v1_xboole_0(u1_struct_0(X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_52, plain, (k2_tarski(X2,X2)=k6_domain_1(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_46, c_0_44])).
cnf(c_0_53, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_54, negated_conjecture, (X1=k1_yellow_0(esk1_0,X2)|~r2_lattice3(esk1_0,X2,X1)|~r1_yellow_0(esk1_0,X2)|~r1_orders_2(esk1_0,X1,k1_yellow_0(esk1_0,X2))|~m1_subset_1(k1_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_23])])).
cnf(c_0_55, plain, (r1_orders_2(X1,X2,k1_yellow_0(X1,k2_tarski(X2,X2)))|~r1_yellow_0(X1,k2_tarski(X2,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_42])).
cnf(c_0_56, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_57, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_52, c_0_53])).
fof(c_0_58, plain, ![X19]:(~l1_orders_2(X19)|l1_struct_0(X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_59, negated_conjecture, (k1_yellow_0(esk1_0,k2_tarski(X1,X1))=X1|~r2_lattice3(esk1_0,k2_tarski(X1,X1),X1)|~r1_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(k1_yellow_0(esk1_0,k2_tarski(X1,X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_23])])).
cnf(c_0_60, plain, (r2_lattice3(X1,k1_tarski(X2),X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_61, plain, ![X43, X44]:((r1_yellow_0(X43,k6_domain_1(u1_struct_0(X43),X44))|~m1_subset_1(X44,u1_struct_0(X43))|(v2_struct_0(X43)|(~v3_orders_2(X43)|(~v5_orders_2(X43)|~l1_orders_2(X43)))))&(r2_yellow_0(X43,k6_domain_1(u1_struct_0(X43),X44))|~m1_subset_1(X44,u1_struct_0(X43))|(v2_struct_0(X43)|(~v3_orders_2(X43)|(~v5_orders_2(X43)|~l1_orders_2(X43)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t38_yellow_0])])])])])).
cnf(c_0_62, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_24])).
cnf(c_0_63, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_64, negated_conjecture, (k1_yellow_0(esk1_0,k2_tarski(X1,X1))=X1|~r2_lattice3(esk1_0,k2_tarski(X1,X1),X1)|~r1_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_42]), c_0_23])])).
cnf(c_0_65, plain, (r2_lattice3(X1,k2_tarski(X2,X2),X3)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[c_0_60, c_0_44])).
cnf(c_0_66, plain, (r1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_67, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_23])])).
fof(c_0_68, plain, ![X8, X9, X11, X12, X14, X15, X18]:(((m1_subset_1(esk3_2(X8,X9),u1_struct_0(X8))|~r2_yellow_0(X8,X9)|~l1_orders_2(X8))&((r1_lattice3(X8,X9,esk3_2(X8,X9))|~r2_yellow_0(X8,X9)|~l1_orders_2(X8))&((~m1_subset_1(X11,u1_struct_0(X8))|(~r1_lattice3(X8,X9,X11)|r1_orders_2(X8,X11,esk3_2(X8,X9)))|~r2_yellow_0(X8,X9)|~l1_orders_2(X8))&((m1_subset_1(esk4_3(X8,X9,X12),u1_struct_0(X8))|~r1_lattice3(X8,X9,X12)|X12=esk3_2(X8,X9)|~m1_subset_1(X12,u1_struct_0(X8))|~r2_yellow_0(X8,X9)|~l1_orders_2(X8))&((r1_lattice3(X8,X9,esk4_3(X8,X9,X12))|~r1_lattice3(X8,X9,X12)|X12=esk3_2(X8,X9)|~m1_subset_1(X12,u1_struct_0(X8))|~r2_yellow_0(X8,X9)|~l1_orders_2(X8))&(~r1_orders_2(X8,esk4_3(X8,X9,X12),X12)|~r1_lattice3(X8,X9,X12)|X12=esk3_2(X8,X9)|~m1_subset_1(X12,u1_struct_0(X8))|~r2_yellow_0(X8,X9)|~l1_orders_2(X8)))))))&(((m1_subset_1(esk6_3(X8,X14,X15),u1_struct_0(X8))|m1_subset_1(esk5_3(X8,X14,X15),u1_struct_0(X8))|~r1_lattice3(X8,X14,X15)|~m1_subset_1(X15,u1_struct_0(X8))|r2_yellow_0(X8,X14)|~l1_orders_2(X8))&(((r1_lattice3(X8,X14,esk6_3(X8,X14,X15))|m1_subset_1(esk5_3(X8,X14,X15),u1_struct_0(X8))|~r1_lattice3(X8,X14,X15)|~m1_subset_1(X15,u1_struct_0(X8))|r2_yellow_0(X8,X14)|~l1_orders_2(X8))&(~m1_subset_1(X18,u1_struct_0(X8))|(~r1_lattice3(X8,X14,X18)|r1_orders_2(X8,X18,esk6_3(X8,X14,X15)))|m1_subset_1(esk5_3(X8,X14,X15),u1_struct_0(X8))|~r1_lattice3(X8,X14,X15)|~m1_subset_1(X15,u1_struct_0(X8))|r2_yellow_0(X8,X14)|~l1_orders_2(X8)))&(esk6_3(X8,X14,X15)!=X15|m1_subset_1(esk5_3(X8,X14,X15),u1_struct_0(X8))|~r1_lattice3(X8,X14,X15)|~m1_subset_1(X15,u1_struct_0(X8))|r2_yellow_0(X8,X14)|~l1_orders_2(X8))))&(((m1_subset_1(esk6_3(X8,X14,X15),u1_struct_0(X8))|r1_lattice3(X8,X14,esk5_3(X8,X14,X15))|~r1_lattice3(X8,X14,X15)|~m1_subset_1(X15,u1_struct_0(X8))|r2_yellow_0(X8,X14)|~l1_orders_2(X8))&(((r1_lattice3(X8,X14,esk6_3(X8,X14,X15))|r1_lattice3(X8,X14,esk5_3(X8,X14,X15))|~r1_lattice3(X8,X14,X15)|~m1_subset_1(X15,u1_struct_0(X8))|r2_yellow_0(X8,X14)|~l1_orders_2(X8))&(~m1_subset_1(X18,u1_struct_0(X8))|(~r1_lattice3(X8,X14,X18)|r1_orders_2(X8,X18,esk6_3(X8,X14,X15)))|r1_lattice3(X8,X14,esk5_3(X8,X14,X15))|~r1_lattice3(X8,X14,X15)|~m1_subset_1(X15,u1_struct_0(X8))|r2_yellow_0(X8,X14)|~l1_orders_2(X8)))&(esk6_3(X8,X14,X15)!=X15|r1_lattice3(X8,X14,esk5_3(X8,X14,X15))|~r1_lattice3(X8,X14,X15)|~m1_subset_1(X15,u1_struct_0(X8))|r2_yellow_0(X8,X14)|~l1_orders_2(X8))))&((m1_subset_1(esk6_3(X8,X14,X15),u1_struct_0(X8))|~r1_orders_2(X8,esk5_3(X8,X14,X15),X15)|~r1_lattice3(X8,X14,X15)|~m1_subset_1(X15,u1_struct_0(X8))|r2_yellow_0(X8,X14)|~l1_orders_2(X8))&(((r1_lattice3(X8,X14,esk6_3(X8,X14,X15))|~r1_orders_2(X8,esk5_3(X8,X14,X15),X15)|~r1_lattice3(X8,X14,X15)|~m1_subset_1(X15,u1_struct_0(X8))|r2_yellow_0(X8,X14)|~l1_orders_2(X8))&(~m1_subset_1(X18,u1_struct_0(X8))|(~r1_lattice3(X8,X14,X18)|r1_orders_2(X8,X18,esk6_3(X8,X14,X15)))|~r1_orders_2(X8,esk5_3(X8,X14,X15),X15)|~r1_lattice3(X8,X14,X15)|~m1_subset_1(X15,u1_struct_0(X8))|r2_yellow_0(X8,X14)|~l1_orders_2(X8)))&(esk6_3(X8,X14,X15)!=X15|~r1_orders_2(X8,esk5_3(X8,X14,X15),X15)|~r1_lattice3(X8,X14,X15)|~m1_subset_1(X15,u1_struct_0(X8))|r2_yellow_0(X8,X14)|~l1_orders_2(X8))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_yellow_0])])])])])])).
cnf(c_0_69, plain, (r1_orders_2(X1,X3,X2)|~r1_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_70, negated_conjecture, (k1_yellow_0(esk1_0,k2_tarski(X1,X1))=X1|~r1_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_23])]), c_0_27])).
cnf(c_0_71, negated_conjecture, (r1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_53]), c_0_23]), c_0_40]), c_0_22])]), c_0_24])).
cnf(c_0_72, plain, (r1_orders_2(X2,X1,esk3_2(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~r2_yellow_0(X2,X3)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_73, plain, (r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,k2_tarski(X2,X2),X3)), inference(rw,[status(thm)],[c_0_69, c_0_44])).
cnf(c_0_74, plain, (r1_lattice3(X1,X2,esk3_2(X1,X2))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_75, plain, (m1_subset_1(esk3_2(X1,X2),u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_76, negated_conjecture, (k1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_53])])).
cnf(c_0_77, negated_conjecture, (esk3_2(esk1_0,X1)=X2|~r1_orders_2(esk1_0,esk3_2(esk1_0,X1),X2)|~r1_lattice3(esk1_0,X1,X2)|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(esk3_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_72]), c_0_23])])).
cnf(c_0_78, plain, (r1_orders_2(X1,esk3_2(X1,k2_tarski(X2,X2)),X2)|~r2_yellow_0(X1,k2_tarski(X2,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_75])).
cnf(c_0_79, negated_conjecture, (X1=esk2_0|~r2_lattice3(esk1_0,k2_tarski(esk2_0,esk2_0),X1)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_76]), c_0_71]), c_0_53])])).
cnf(c_0_80, plain, (r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_81, negated_conjecture, (esk3_2(esk1_0,k2_tarski(X1,X1))=X1|~r1_lattice3(esk1_0,k2_tarski(X1,X1),X1)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(esk3_2(esk1_0,k2_tarski(X1,X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_23])])).
cnf(c_0_82, plain, (r1_lattice3(X1,k1_tarski(X3),X2)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_83, negated_conjecture, (X1=esk2_0|~r1_orders_2(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_65]), c_0_53]), c_0_23])])).
cnf(c_0_84, negated_conjecture, (r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_67]), c_0_53]), c_0_23]), c_0_40]), c_0_22])]), c_0_24])).
cnf(c_0_85, negated_conjecture, (esk3_2(esk1_0,k2_tarski(X1,X1))=X1|~r1_lattice3(esk1_0,k2_tarski(X1,X1),X1)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_75]), c_0_23])])).
cnf(c_0_86, plain, (r1_lattice3(X1,k2_tarski(X3,X3),X2)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[c_0_82, c_0_44])).
cnf(c_0_87, negated_conjecture, (esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0))=esk2_0|~r1_orders_2(esk1_0,esk2_0,esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_78]), c_0_84]), c_0_53]), c_0_23])])).
cnf(c_0_88, negated_conjecture, (esk3_2(esk1_0,k2_tarski(X1,X1))=X1|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_23])]), c_0_27])).
fof(c_0_89, plain, ![X35, X36, X37, X38, X39]:(((r1_lattice3(X35,X37,X36)|(X36!=k2_yellow_0(X35,X37)|~r2_yellow_0(X35,X37))|~m1_subset_1(X36,u1_struct_0(X35))|(~v5_orders_2(X35)|~l1_orders_2(X35)))&(~m1_subset_1(X38,u1_struct_0(X35))|(~r1_lattice3(X35,X37,X38)|r1_orders_2(X35,X38,X36))|(X36!=k2_yellow_0(X35,X37)|~r2_yellow_0(X35,X37))|~m1_subset_1(X36,u1_struct_0(X35))|(~v5_orders_2(X35)|~l1_orders_2(X35))))&(((X36=k2_yellow_0(X35,X39)|(m1_subset_1(esk8_3(X35,X36,X39),u1_struct_0(X35))|~r1_lattice3(X35,X39,X36))|~m1_subset_1(X36,u1_struct_0(X35))|(~v5_orders_2(X35)|~l1_orders_2(X35)))&(r2_yellow_0(X35,X39)|(m1_subset_1(esk8_3(X35,X36,X39),u1_struct_0(X35))|~r1_lattice3(X35,X39,X36))|~m1_subset_1(X36,u1_struct_0(X35))|(~v5_orders_2(X35)|~l1_orders_2(X35))))&(((X36=k2_yellow_0(X35,X39)|(r1_lattice3(X35,X39,esk8_3(X35,X36,X39))|~r1_lattice3(X35,X39,X36))|~m1_subset_1(X36,u1_struct_0(X35))|(~v5_orders_2(X35)|~l1_orders_2(X35)))&(r2_yellow_0(X35,X39)|(r1_lattice3(X35,X39,esk8_3(X35,X36,X39))|~r1_lattice3(X35,X39,X36))|~m1_subset_1(X36,u1_struct_0(X35))|(~v5_orders_2(X35)|~l1_orders_2(X35))))&((X36=k2_yellow_0(X35,X39)|(~r1_orders_2(X35,esk8_3(X35,X36,X39),X36)|~r1_lattice3(X35,X39,X36))|~m1_subset_1(X36,u1_struct_0(X35))|(~v5_orders_2(X35)|~l1_orders_2(X35)))&(r2_yellow_0(X35,X39)|(~r1_orders_2(X35,esk8_3(X35,X36,X39),X36)|~r1_lattice3(X35,X39,X36))|~m1_subset_1(X36,u1_struct_0(X35))|(~v5_orders_2(X35)|~l1_orders_2(X35))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_yellow_0])])])])])])).
cnf(c_0_90, negated_conjecture, (esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0))=esk2_0|~r1_lattice3(esk1_0,k2_tarski(esk2_0,esk2_0),esk2_0)|~m1_subset_1(esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_72]), c_0_84]), c_0_53]), c_0_23])])).
cnf(c_0_91, negated_conjecture, (r1_lattice3(esk1_0,k2_tarski(X1,X1),X1)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_88]), c_0_23])])).
cnf(c_0_92, plain, (X1=k2_yellow_0(X2,X3)|~r1_orders_2(X2,esk8_3(X2,X1,X3),X1)|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_93, negated_conjecture, (esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0))=esk2_0|~r1_lattice3(esk1_0,k2_tarski(esk2_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_75]), c_0_84]), c_0_23])])).
cnf(c_0_94, negated_conjecture, (r1_lattice3(esk1_0,k2_tarski(esk2_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_84]), c_0_53])])).
cnf(c_0_95, negated_conjecture, (k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_96, plain, (esk3_2(X1,X2)=k2_yellow_0(X1,X3)|~r1_lattice3(X1,X2,esk8_3(X1,esk3_2(X1,X2),X3))|~r1_lattice3(X1,X3,esk3_2(X1,X2))|~r2_yellow_0(X1,X2)|~m1_subset_1(esk8_3(X1,esk3_2(X1,X2),X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_72]), c_0_75])).
cnf(c_0_97, negated_conjecture, (esk3_2(esk1_0,k2_tarski(esk2_0,esk2_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_94])])).
cnf(c_0_98, negated_conjecture, (k1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))!=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_95, c_0_67]), c_0_67])).
cnf(c_0_99, negated_conjecture, (k2_yellow_0(esk1_0,X1)=esk2_0|~r1_lattice3(esk1_0,k2_tarski(esk2_0,esk2_0),esk8_3(esk1_0,esk2_0,X1))|~r1_lattice3(esk1_0,X1,esk2_0)|~m1_subset_1(esk8_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_84]), c_0_23]), c_0_40])])).
cnf(c_0_100, plain, (X1=k2_yellow_0(X2,X3)|r1_lattice3(X2,X3,esk8_3(X2,X1,X3))|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_101, negated_conjecture, (k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_98, c_0_76])])).
cnf(c_0_102, negated_conjecture, (~m1_subset_1(esk8_3(esk1_0,esk2_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_94]), c_0_53]), c_0_23]), c_0_40])]), c_0_101])).
cnf(c_0_103, plain, (X1=k2_yellow_0(X2,X3)|m1_subset_1(esk8_3(X2,X1,X3),u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_104, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_94]), c_0_53]), c_0_23]), c_0_40])]), c_0_101]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 105
# Proof object clause steps            : 74
# Proof object formula steps           : 31
# Proof object conjectures             : 41
# Proof object clause conjectures      : 38
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 30
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 34
# Proof object simplifying inferences  : 96
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 61
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 60
# Processed clauses                    : 228
# ...of these trivial                  : 1
# ...subsumed                          : 27
# ...remaining for further processing  : 200
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 9
# Backward-rewritten                   : 8
# Generated clauses                    : 351
# ...of the previous two non-trivial   : 288
# Contextual simplify-reflections      : 11
# Paramodulations                      : 344
# Factorizations                       : 0
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 116
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 100
# Current number of unprocessed clauses: 180
# ...number of literals in the above   : 1440
# Current number of archived formulas  : 0
# Current number of archived clauses   : 78
# Clause-clause subsumption calls (NU) : 4473
# Rec. Clause-clause subsumption calls : 468
# Non-unit clause-clause subsumptions  : 47
# Unit Clause-clause subsumption calls : 21
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 17470

# -------------------------------------------------
# User time                : 0.018 s
# System time              : 0.003 s
# Total time               : 0.022 s
# Maximum resident set size: 28608 pages
