# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_yellow_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>(v2_lattice3(k2_yellow_1(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>r1_tarski(k11_lattice3(k2_yellow_1(X1),X2,X3),k3_xboole_0(X2,X3)))))), file('data/MPTP/proofs/64cf61cc192d4b4db0e698e512822872.E_input', t7_yellow_1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/64cf61cc192d4b4db0e698e512822872.E_input', t12_setfam_1)).
fof(t19_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X1,X3))=>r1_tarski(X1,k3_xboole_0(X2,X3))), file('data/MPTP/proofs/64cf61cc192d4b4db0e698e512822872.E_input', t19_xboole_1)).
fof(t3_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r3_orders_2(k2_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3))))), file('data/MPTP/proofs/64cf61cc192d4b4db0e698e512822872.E_input', t3_yellow_1)).
fof(t1_yellow_1, axiom, ![X1]:(u1_struct_0(k2_yellow_1(X1))=X1&u1_orders_2(k2_yellow_1(X1))=k1_yellow_1(X1)), file('data/MPTP/proofs/64cf61cc192d4b4db0e698e512822872.E_input', t1_yellow_1)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/64cf61cc192d4b4db0e698e512822872.E_input', redefinition_r3_orders_2)).
fof(dt_k2_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&l1_orders_2(k2_yellow_1(X1))), file('data/MPTP/proofs/64cf61cc192d4b4db0e698e512822872.E_input', dt_k2_yellow_1)).
fof(fc5_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&(v3_orders_2(k2_yellow_1(X1))&(v4_orders_2(k2_yellow_1(X1))&v5_orders_2(k2_yellow_1(X1))))), file('data/MPTP/proofs/64cf61cc192d4b4db0e698e512822872.E_input', fc5_yellow_1)).
fof(fc6_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>(~(v2_struct_0(k2_yellow_1(X1)))&v1_orders_2(k2_yellow_1(X1)))), file('data/MPTP/proofs/64cf61cc192d4b4db0e698e512822872.E_input', fc6_yellow_1)).
fof(t23_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k12_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))))))), file('data/MPTP/proofs/64cf61cc192d4b4db0e698e512822872.E_input', t23_yellow_0)).
fof(dt_k12_lattice3, axiom, ![X1, X2, X3]:(((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k12_lattice3(X1,X2,X3),u1_struct_0(X1))), file('data/MPTP/proofs/64cf61cc192d4b4db0e698e512822872.E_input', dt_k12_lattice3)).
fof(redefinition_k12_lattice3, axiom, ![X1, X2, X3]:(((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)), file('data/MPTP/proofs/64cf61cc192d4b4db0e698e512822872.E_input', redefinition_k12_lattice3)).
fof(c_0_12, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>(v2_lattice3(k2_yellow_1(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>r1_tarski(k11_lattice3(k2_yellow_1(X1),X2,X3),k3_xboole_0(X2,X3))))))), inference(assume_negation,[status(cth)],[t7_yellow_1])).
fof(c_0_13, negated_conjecture, (~v1_xboole_0(esk1_0)&(v2_lattice3(k2_yellow_1(esk1_0))&(m1_subset_1(esk2_0,u1_struct_0(k2_yellow_1(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(k2_yellow_1(esk1_0)))&~r1_tarski(k11_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0),k3_xboole_0(esk2_0,esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])).
fof(c_0_14, plain, ![X9, X10]:k1_setfam_1(k2_tarski(X9,X10))=k3_xboole_0(X9,X10), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_15, plain, ![X11, X12, X13]:(~r1_tarski(X11,X12)|~r1_tarski(X11,X13)|r1_tarski(X11,k3_xboole_0(X12,X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t19_xboole_1])])).
fof(c_0_16, plain, ![X14, X15, X16]:((~r3_orders_2(k2_yellow_1(X14),X15,X16)|r1_tarski(X15,X16)|~m1_subset_1(X16,u1_struct_0(k2_yellow_1(X14)))|~m1_subset_1(X15,u1_struct_0(k2_yellow_1(X14)))|v1_xboole_0(X14))&(~r1_tarski(X15,X16)|r3_orders_2(k2_yellow_1(X14),X15,X16)|~m1_subset_1(X16,u1_struct_0(k2_yellow_1(X14)))|~m1_subset_1(X15,u1_struct_0(k2_yellow_1(X14)))|v1_xboole_0(X14))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_yellow_1])])])])])).
fof(c_0_17, plain, ![X17]:(u1_struct_0(k2_yellow_1(X17))=X17&u1_orders_2(k2_yellow_1(X17))=k1_yellow_1(X17)), inference(variable_rename,[status(thm)],[t1_yellow_1])).
cnf(c_0_18, negated_conjecture, (~r1_tarski(k11_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0),k3_xboole_0(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (r1_tarski(X1,k3_xboole_0(X2,X3))|~r1_tarski(X1,X2)|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (r1_tarski(X2,X3)|v1_xboole_0(X1)|~r3_orders_2(k2_yellow_1(X1),X2,X3)|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (u1_struct_0(k2_yellow_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_23, plain, ![X18, X19, X20]:((~r3_orders_2(X18,X19,X20)|r1_orders_2(X18,X19,X20)|(v2_struct_0(X18)|(~v3_orders_2(X18)|~l1_orders_2(X18))|(~m1_subset_1(X19,u1_struct_0(X18))|~m1_subset_1(X20,u1_struct_0(X18)))))&(~r1_orders_2(X18,X19,X20)|r3_orders_2(X18,X19,X20)|(v2_struct_0(X18)|(~v3_orders_2(X18)|~l1_orders_2(X18))|(~m1_subset_1(X19,u1_struct_0(X18))|~m1_subset_1(X20,u1_struct_0(X18)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
fof(c_0_24, plain, ![X21]:(v1_orders_2(k2_yellow_1(X21))&l1_orders_2(k2_yellow_1(X21))), inference(variable_rename,[status(thm)],[dt_k2_yellow_1])).
fof(c_0_25, plain, ![X22]:(v1_orders_2(k2_yellow_1(X22))&(v3_orders_2(k2_yellow_1(X22))&(v4_orders_2(k2_yellow_1(X22))&v5_orders_2(k2_yellow_1(X22))))), inference(variable_rename,[status(thm)],[fc5_yellow_1])).
fof(c_0_26, plain, ![X23]:((~v2_struct_0(k2_yellow_1(X23))|v1_xboole_0(X23))&(v1_orders_2(k2_yellow_1(X23))|v1_xboole_0(X23))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc6_yellow_1])])])])).
fof(c_0_27, plain, ![X24, X25, X26, X27, X28]:(((r1_orders_2(X24,X27,X25)|X27!=k12_lattice3(X24,X25,X26)|~m1_subset_1(X27,u1_struct_0(X24))|~m1_subset_1(X26,u1_struct_0(X24))|~m1_subset_1(X25,u1_struct_0(X24))|(~v5_orders_2(X24)|(~v2_lattice3(X24)|~l1_orders_2(X24))))&((r1_orders_2(X24,X27,X26)|X27!=k12_lattice3(X24,X25,X26)|~m1_subset_1(X27,u1_struct_0(X24))|~m1_subset_1(X26,u1_struct_0(X24))|~m1_subset_1(X25,u1_struct_0(X24))|(~v5_orders_2(X24)|(~v2_lattice3(X24)|~l1_orders_2(X24))))&(~m1_subset_1(X28,u1_struct_0(X24))|(~r1_orders_2(X24,X28,X25)|~r1_orders_2(X24,X28,X26)|r1_orders_2(X24,X28,X27))|X27!=k12_lattice3(X24,X25,X26)|~m1_subset_1(X27,u1_struct_0(X24))|~m1_subset_1(X26,u1_struct_0(X24))|~m1_subset_1(X25,u1_struct_0(X24))|(~v5_orders_2(X24)|(~v2_lattice3(X24)|~l1_orders_2(X24))))))&((m1_subset_1(esk4_4(X24,X25,X26,X27),u1_struct_0(X24))|~r1_orders_2(X24,X27,X26)|~r1_orders_2(X24,X27,X25)|X27=k12_lattice3(X24,X25,X26)|~m1_subset_1(X27,u1_struct_0(X24))|~m1_subset_1(X26,u1_struct_0(X24))|~m1_subset_1(X25,u1_struct_0(X24))|(~v5_orders_2(X24)|(~v2_lattice3(X24)|~l1_orders_2(X24))))&(((r1_orders_2(X24,esk4_4(X24,X25,X26,X27),X25)|~r1_orders_2(X24,X27,X26)|~r1_orders_2(X24,X27,X25)|X27=k12_lattice3(X24,X25,X26)|~m1_subset_1(X27,u1_struct_0(X24))|~m1_subset_1(X26,u1_struct_0(X24))|~m1_subset_1(X25,u1_struct_0(X24))|(~v5_orders_2(X24)|(~v2_lattice3(X24)|~l1_orders_2(X24))))&(r1_orders_2(X24,esk4_4(X24,X25,X26,X27),X26)|~r1_orders_2(X24,X27,X26)|~r1_orders_2(X24,X27,X25)|X27=k12_lattice3(X24,X25,X26)|~m1_subset_1(X27,u1_struct_0(X24))|~m1_subset_1(X26,u1_struct_0(X24))|~m1_subset_1(X25,u1_struct_0(X24))|(~v5_orders_2(X24)|(~v2_lattice3(X24)|~l1_orders_2(X24)))))&(~r1_orders_2(X24,esk4_4(X24,X25,X26,X27),X27)|~r1_orders_2(X24,X27,X26)|~r1_orders_2(X24,X27,X25)|X27=k12_lattice3(X24,X25,X26)|~m1_subset_1(X27,u1_struct_0(X24))|~m1_subset_1(X26,u1_struct_0(X24))|~m1_subset_1(X25,u1_struct_0(X24))|(~v5_orders_2(X24)|(~v2_lattice3(X24)|~l1_orders_2(X24))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_yellow_0])])])])])).
fof(c_0_28, plain, ![X30, X31, X32]:(~v5_orders_2(X30)|(~v2_lattice3(X30)|~l1_orders_2(X30))|(~m1_subset_1(X31,u1_struct_0(X30))|~m1_subset_1(X32,u1_struct_0(X30)))|m1_subset_1(k12_lattice3(X30,X31,X32),u1_struct_0(X30))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k12_lattice3])])).
cnf(c_0_29, negated_conjecture, (~r1_tarski(k11_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0),k1_setfam_1(k2_tarski(esk2_0,esk3_0)))), inference(rw,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_30, plain, (r1_tarski(X1,k1_setfam_1(k2_tarski(X2,X3)))|~r1_tarski(X1,X3)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_20, c_0_19])).
fof(c_0_31, plain, ![X33, X34, X35]:(~v5_orders_2(X33)|(~v2_lattice3(X33)|~l1_orders_2(X33))|(~m1_subset_1(X34,u1_struct_0(X33))|~m1_subset_1(X35,u1_struct_0(X33)))|k12_lattice3(X33,X34,X35)=k11_lattice3(X33,X34,X35)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k12_lattice3])])).
cnf(c_0_32, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k2_yellow_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_33, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(k2_yellow_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_34, plain, (r1_tarski(X1,X2)|v1_xboole_0(X3)|~r3_orders_2(k2_yellow_1(X3),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_22]), c_0_22])).
cnf(c_0_35, plain, (r3_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_36, plain, (l1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_37, plain, (v3_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_38, plain, (v1_xboole_0(X1)|~v2_struct_0(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_39, plain, (r1_orders_2(X1,X2,X3)|X2!=k12_lattice3(X1,X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_40, plain, (m1_subset_1(k12_lattice3(X1,X2,X3),u1_struct_0(X1))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_41, plain, (v5_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_42, negated_conjecture, (~r1_tarski(k11_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0),esk3_0)|~r1_tarski(k11_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0),esk2_0)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_43, plain, (k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_44, negated_conjecture, (m1_subset_1(esk3_0,esk1_0)), inference(rw,[status(thm)],[c_0_32, c_0_22])).
cnf(c_0_45, negated_conjecture, (m1_subset_1(esk2_0,esk1_0)), inference(rw,[status(thm)],[c_0_33, c_0_22])).
cnf(c_0_46, negated_conjecture, (v2_lattice3(k2_yellow_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_47, plain, (r1_tarski(X1,X2)|v1_xboole_0(X3)|~r1_orders_2(k2_yellow_1(X3),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_37]), c_0_22]), c_0_22])]), c_0_38])).
cnf(c_0_48, plain, (r1_orders_2(X1,k12_lattice3(X1,X2,X3),X3)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v2_lattice3(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_39]), c_0_40])).
cnf(c_0_49, plain, (m1_subset_1(k12_lattice3(k2_yellow_1(X1),X2,X3),X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v2_lattice3(k2_yellow_1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_22]), c_0_41]), c_0_36])])).
cnf(c_0_50, plain, (r1_orders_2(X1,X2,X3)|X2!=k12_lattice3(X1,X3,X4)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_51, negated_conjecture, (~r1_tarski(k12_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0),esk3_0)|~r1_tarski(k12_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_41]), c_0_36]), c_0_22]), c_0_44]), c_0_22]), c_0_45]), c_0_46])])).
cnf(c_0_52, plain, (r1_tarski(k12_lattice3(k2_yellow_1(X1),X2,X3),X3)|v1_xboole_0(X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v2_lattice3(k2_yellow_1(X1))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_41]), c_0_36]), c_0_22]), c_0_22])]), c_0_49])).
cnf(c_0_53, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_54, plain, (r1_orders_2(X1,k12_lattice3(X1,X2,X3),X2)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_lattice3(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_50]), c_0_40])).
cnf(c_0_55, negated_conjecture, (~r1_tarski(k12_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0),esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_44]), c_0_45]), c_0_46])]), c_0_53])).
cnf(c_0_56, plain, (r1_tarski(k12_lattice3(k2_yellow_1(X1),X2,X3),X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)|~v2_lattice3(k2_yellow_1(X1))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_54]), c_0_41]), c_0_36]), c_0_22]), c_0_22])]), c_0_49])).
cnf(c_0_57, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_45]), c_0_44]), c_0_46])]), c_0_53]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 58
# Proof object clause steps            : 33
# Proof object formula steps           : 25
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 49
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 28
# Processed clauses                    : 40
# ...of these trivial                  : 2
# ...subsumed                          : 3
# ...remaining for further processing  : 35
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 21
# ...of the previous two non-trivial   : 18
# Contextual simplify-reflections      : 6
# Paramodulations                      : 19
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 33
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 20
# Current number of unprocessed clauses: 5
# ...number of literals in the above   : 52
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2
# Clause-clause subsumption calls (NU) : 156
# Rec. Clause-clause subsumption calls : 12
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2953

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28744 pages
