# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', t3_xboole_0)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', t65_xboole_1)).
fof(t7_tarski, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&r2_hidden(X4,X3))))))), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', t7_tarski)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', d1_tarski)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', d4_relat_1)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', cc2_relat_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', t3_subset)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', t36_xboole_1)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', redefinition_k6_subset_1)).
fof(t43_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3)), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', t43_xboole_1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', l51_zfmisc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', d10_xboole_0)).
fof(t2_tarski, axiom, ![X1, X2]:(![X3]:(r2_hidden(X3,X1)<=>r2_hidden(X3,X2))=>X1=X2), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', t2_tarski)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', t2_xboole_1)).
fof(t31_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>r1_tarski(k3_relat_1(X1),k3_relat_1(X2))))), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', t31_relat_1)).
fof(t10_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(X1,k2_xboole_0(X3,X2))), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', t10_xboole_1)).
fof(t8_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X3,X2))=>r1_tarski(k2_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', t8_xboole_1)).
fof(t25_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>(r1_tarski(k1_relat_1(X1),k1_relat_1(X2))&r1_tarski(k2_relat_1(X1),k2_relat_1(X2)))))), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', t25_relat_1)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', t1_xboole_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', commutativity_k2_tarski)).
fof(t44_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k4_xboole_0(X1,X2),X3)=>r1_tarski(X1,k2_xboole_0(X2,X3))), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', t44_xboole_1)).
fof(d6_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k3_relat_1(X1)=k2_xboole_0(k1_relat_1(X1),k2_relat_1(X1))), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', d6_relat_1)).
fof(t15_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>r1_tarski(k6_subset_1(k1_relat_1(X1),k1_relat_1(X2)),k1_relat_1(k6_subset_1(X1,X2))))), file('data/MPTP/proofs/64db2f91267d4060badf8fcb82a4f2c9.E_input', t15_relat_1)).
fof(c_0_23, plain, ![X42, X43, X45, X46, X47]:(((r2_hidden(esk7_2(X42,X43),X42)|r1_xboole_0(X42,X43))&(r2_hidden(esk7_2(X42,X43),X43)|r1_xboole_0(X42,X43)))&(~r2_hidden(X47,X45)|~r2_hidden(X47,X46)|~r1_xboole_0(X45,X46))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_24, plain, ![X41]:r1_xboole_0(X41,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_25, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_26, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_27, plain, ![X11, X12, X14]:((r2_hidden(esk3_2(X11,X12),X12)|~r2_hidden(X11,X12))&(~r2_hidden(X14,X12)|~r2_hidden(X14,esk3_2(X11,X12))|~r2_hidden(X11,X12))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tarski])])])])])).
fof(c_0_28, plain, ![X67, X68, X69, X70, X71, X72]:(((~r2_hidden(X69,X68)|X69=X67|X68!=k1_tarski(X67))&(X70!=X67|r2_hidden(X70,X68)|X68!=k1_tarski(X67)))&((~r2_hidden(esk9_2(X71,X72),X72)|esk9_2(X71,X72)!=X71|X72=k1_tarski(X71))&(r2_hidden(esk9_2(X71,X72),X72)|esk9_2(X71,X72)=X71|X72=k1_tarski(X71)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_29, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_30, plain, (r2_hidden(esk3_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_31, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_32, plain, ![X19, X20, X21, X23, X24, X25, X26, X28]:(((~r2_hidden(X21,X20)|r2_hidden(k4_tarski(X21,esk4_3(X19,X20,X21)),X19)|X20!=k1_relat_1(X19))&(~r2_hidden(k4_tarski(X23,X24),X19)|r2_hidden(X23,X20)|X20!=k1_relat_1(X19)))&((~r2_hidden(esk5_2(X25,X26),X26)|~r2_hidden(k4_tarski(esk5_2(X25,X26),X28),X25)|X26=k1_relat_1(X25))&(r2_hidden(esk5_2(X25,X26),X26)|r2_hidden(k4_tarski(esk5_2(X25,X26),esk6_2(X25,X26)),X25)|X26=k1_relat_1(X25)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
fof(c_0_33, plain, ![X51, X52]:(~v1_relat_1(X51)|(~m1_subset_1(X52,k1_zfmisc_1(X51))|v1_relat_1(X52))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_34, plain, ![X7, X8]:((~m1_subset_1(X7,k1_zfmisc_1(X8))|r1_tarski(X7,X8))&(~r1_tarski(X7,X8)|m1_subset_1(X7,k1_zfmisc_1(X8)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_35, plain, ![X65, X66]:r1_tarski(k4_xboole_0(X65,X66),X65), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_36, plain, ![X30, X31]:k6_subset_1(X30,X31)=k4_xboole_0(X30,X31), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
fof(c_0_37, plain, ![X56, X57, X58]:(~r1_tarski(X56,k2_xboole_0(X57,X58))|r1_tarski(k4_xboole_0(X56,X57),X58)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
fof(c_0_38, plain, ![X48, X49]:k3_tarski(k2_tarski(X48,X49))=k2_xboole_0(X48,X49), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_39, plain, ![X15, X16]:(((r1_tarski(X15,X16)|X15!=X16)&(r1_tarski(X16,X15)|X15!=X16))&(~r1_tarski(X15,X16)|~r1_tarski(X16,X15)|X15=X16)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_40, plain, (~r2_hidden(esk3_2(X1,k1_xboole_0),X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_41, plain, (r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_31])])).
cnf(c_0_42, plain, (r2_hidden(k4_tarski(X1,esk4_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_43, plain, ![X53, X54]:((~r2_hidden(esk8_2(X53,X54),X53)|~r2_hidden(esk8_2(X53,X54),X54)|X53=X54)&(r2_hidden(esk8_2(X53,X54),X53)|r2_hidden(esk8_2(X53,X54),X54)|X53=X54)), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tarski])])])])).
cnf(c_0_44, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_45, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_46, plain, ![X50]:r1_tarski(k1_xboole_0,X50), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_47, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>r1_tarski(k3_relat_1(X1),k3_relat_1(X2)))))), inference(assume_negation,[status(cth)],[t31_relat_1])).
fof(c_0_48, plain, ![X59, X60, X61]:(~r1_tarski(X59,X60)|r1_tarski(X59,k2_xboole_0(X61,X60))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_xboole_1])])).
cnf(c_0_49, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_50, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_51, plain, ![X32, X33, X34]:(~r1_tarski(X32,X33)|~r1_tarski(X34,X33)|r1_tarski(k2_xboole_0(X32,X34),X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_xboole_1])])).
cnf(c_0_52, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_53, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_54, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_55, plain, ![X17, X18]:((r1_tarski(k1_relat_1(X17),k1_relat_1(X18))|~r1_tarski(X17,X18)|~v1_relat_1(X18)|~v1_relat_1(X17))&(r1_tarski(k2_relat_1(X17),k2_relat_1(X18))|~r1_tarski(X17,X18)|~v1_relat_1(X18)|~v1_relat_1(X17))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t25_relat_1])])])])).
cnf(c_0_56, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_57, plain, (r2_hidden(k4_tarski(X1,esk4_3(X2,k1_relat_1(X2),X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_58, plain, (r2_hidden(esk8_2(X1,X2),X1)|r2_hidden(esk8_2(X1,X2),X2)|X1=X2), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_59, plain, (v1_relat_1(X1)|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_60, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_61, negated_conjecture, (v1_relat_1(esk1_0)&(v1_relat_1(esk2_0)&(r1_tarski(esk1_0,esk2_0)&~r1_tarski(k3_relat_1(esk1_0),k3_relat_1(esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_47])])])).
cnf(c_0_62, plain, (r1_tarski(X1,k2_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_63, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_64, plain, (r1_tarski(k6_subset_1(X1,X2),X1)), inference(rw,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_65, plain, (r1_tarski(k2_xboole_0(X1,X3),X2)|~r1_tarski(X1,X2)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_66, plain, (r1_tarski(k6_subset_1(X1,X2),X3)|~r1_tarski(X1,k3_tarski(k2_tarski(X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_50]), c_0_53])).
cnf(c_0_67, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_54])).
fof(c_0_68, plain, ![X38, X39, X40]:(~r1_tarski(X38,X39)|~r1_tarski(X39,X40)|r1_tarski(X38,X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
cnf(c_0_69, plain, (r1_tarski(k1_relat_1(X1),k1_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_70, plain, (~r2_hidden(X1,k1_relat_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_71, plain, (k1_xboole_0=X1|r2_hidden(esk8_2(k1_xboole_0,X1),X1)), inference(spm,[status(thm)],[c_0_56, c_0_58])).
cnf(c_0_72, plain, (v1_relat_1(k1_xboole_0)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_73, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_74, plain, (r1_tarski(X1,k3_tarski(k2_tarski(X3,X2)))|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_62, c_0_53])).
fof(c_0_75, plain, ![X63, X64]:k2_tarski(X63,X64)=k2_tarski(X64,X63), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_76, plain, (k6_subset_1(X1,X2)=X1|~r1_tarski(X1,k6_subset_1(X1,X2))), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_77, plain, (r1_tarski(k3_tarski(k2_tarski(X1,X3)),X2)|~r1_tarski(X3,X2)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_65, c_0_53])).
cnf(c_0_78, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_63, c_0_60])).
cnf(c_0_79, plain, (r1_tarski(k6_subset_1(k3_tarski(k2_tarski(X1,X2)),X1),X2)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_80, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_81, plain, (r1_tarski(k1_relat_1(X1),k1_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(csr,[status(thm)],[c_0_69, c_0_59])).
cnf(c_0_82, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_83, negated_conjecture, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
fof(c_0_84, plain, ![X35, X36, X37]:(~r1_tarski(k4_xboole_0(X35,X36),X37)|r1_tarski(X35,k2_xboole_0(X36,X37))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t44_xboole_1])])).
cnf(c_0_85, plain, (k3_tarski(k2_tarski(X1,X2))=X3|~r1_tarski(k3_tarski(k2_tarski(X1,X2)),X3)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_63, c_0_74])).
cnf(c_0_86, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_87, plain, (k6_subset_1(k3_tarski(k2_tarski(X1,X2)),X3)=k3_tarski(k2_tarski(X1,X2))|~r1_tarski(X2,k6_subset_1(k3_tarski(k2_tarski(X1,X2)),X3))|~r1_tarski(X1,k6_subset_1(k3_tarski(k2_tarski(X1,X2)),X3))), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_88, plain, (k6_subset_1(k3_tarski(k2_tarski(X1,k1_xboole_0)),X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_89, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_80, c_0_60])).
cnf(c_0_90, plain, (r1_tarski(k1_relat_1(X1),k1_xboole_0)|~r1_tarski(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_83])])).
cnf(c_0_91, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(k4_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_92, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_93, plain, (k3_tarski(k2_tarski(X1,X2))=X3|~r1_tarski(k3_tarski(k2_tarski(X1,X2)),X3)|~r1_tarski(X3,X1)), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_94, plain, (k3_tarski(k2_tarski(X1,k1_xboole_0))=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_60])])).
cnf(c_0_95, plain, (r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_96, plain, (r1_tarski(k6_subset_1(X1,X2),X3)|~r1_tarski(X1,k3_tarski(k2_tarski(X3,X2)))), inference(spm,[status(thm)],[c_0_66, c_0_86])).
cnf(c_0_97, plain, (r1_tarski(X1,k3_tarski(k2_tarski(X2,X3)))|~r1_tarski(k6_subset_1(X1,X2),X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_50]), c_0_53])).
cnf(c_0_98, negated_conjecture, (r1_tarski(X1,esk2_0)|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_80, c_0_92])).
cnf(c_0_99, plain, (k1_xboole_0=X1|~r1_tarski(X2,k1_xboole_0)|~r1_tarski(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_60])])).
cnf(c_0_100, plain, (r1_tarski(k6_subset_1(k1_relat_1(X1),X2),X3)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_66, c_0_95])).
cnf(c_0_101, plain, (r1_tarski(k6_subset_1(X1,X2),X3)|~r1_tarski(k6_subset_1(X1,X3),X2)), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_102, negated_conjecture, (r1_tarski(k6_subset_1(esk1_0,X1),esk2_0)), inference(spm,[status(thm)],[c_0_98, c_0_64])).
fof(c_0_103, plain, ![X62]:(~v1_relat_1(X62)|k3_relat_1(X62)=k2_xboole_0(k1_relat_1(X62),k2_relat_1(X62))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_relat_1])])).
cnf(c_0_104, plain, (k1_xboole_0=X1|~r1_tarski(X1,k6_subset_1(k1_relat_1(X2),X3))|~r1_tarski(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_105, negated_conjecture, (r1_tarski(k6_subset_1(esk1_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_106, plain, (r1_tarski(X1,k3_tarski(k2_tarski(X2,X3)))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_74, c_0_86])).
cnf(c_0_107, plain, (k3_relat_1(X1)=k2_xboole_0(k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_103])).
fof(c_0_108, plain, ![X9, X10]:(~v1_relat_1(X9)|(~v1_relat_1(X10)|r1_tarski(k6_subset_1(k1_relat_1(X9),k1_relat_1(X10)),k1_relat_1(k6_subset_1(X9,X10))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_relat_1])])])).
cnf(c_0_109, negated_conjecture, (k6_subset_1(esk1_0,esk2_0)=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_110, plain, (r1_tarski(k6_subset_1(X1,X2),X3)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_66, c_0_106])).
cnf(c_0_111, plain, (k3_relat_1(X1)=k3_tarski(k2_tarski(k1_relat_1(X1),k2_relat_1(X1)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_107, c_0_53])).
cnf(c_0_112, plain, (r1_tarski(k2_relat_1(X1),k2_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_113, plain, (r1_tarski(k6_subset_1(k1_relat_1(X1),k1_relat_1(X2)),k1_relat_1(k6_subset_1(X1,X2)))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_108])).
cnf(c_0_114, negated_conjecture, (k6_subset_1(esk1_0,esk2_0)=k1_xboole_0|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_109, c_0_110])).
cnf(c_0_115, plain, (r1_tarski(X1,k3_relat_1(X2))|~r1_tarski(X1,k2_relat_1(X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_74, c_0_111])).
cnf(c_0_116, plain, (r1_tarski(k2_relat_1(X1),k2_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(csr,[status(thm)],[c_0_112, c_0_59])).
cnf(c_0_117, plain, (k6_subset_1(k1_relat_1(X1),k1_relat_1(X2))=k1_relat_1(k6_subset_1(X1,X2))|~r1_tarski(k1_relat_1(k6_subset_1(X1,X2)),k6_subset_1(k1_relat_1(X1),k1_relat_1(X2)))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_63, c_0_113])).
cnf(c_0_118, negated_conjecture, (k6_subset_1(esk1_0,esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_114, c_0_92])).
cnf(c_0_119, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_120, plain, (r1_tarski(k3_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X2)|~r1_tarski(k1_relat_1(X1),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_77, c_0_111])).
cnf(c_0_121, plain, (r1_tarski(k2_relat_1(X1),k3_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_115, c_0_116])).
cnf(c_0_122, plain, (r1_tarski(X1,k3_relat_1(X2))|~r1_tarski(k6_subset_1(X1,k1_relat_1(X2)),k2_relat_1(X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_97, c_0_111])).
cnf(c_0_123, negated_conjecture, (k6_subset_1(k1_relat_1(esk1_0),k1_relat_1(esk2_0))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_82]), c_0_82]), c_0_60]), c_0_73]), c_0_119])])).
cnf(c_0_124, negated_conjecture, (~r1_tarski(k3_relat_1(esk1_0),k3_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_125, plain, (r1_tarski(k3_relat_1(X1),k3_relat_1(X2))|~r1_tarski(k1_relat_1(X1),k3_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121]), c_0_59])).
cnf(c_0_126, negated_conjecture, (r1_tarski(k1_relat_1(esk1_0),k3_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_60]), c_0_73])])).
cnf(c_0_127, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125]), c_0_126]), c_0_92]), c_0_73])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 128
# Proof object clause steps            : 81
# Proof object formula steps           : 47
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 23
# Proof object generating inferences   : 42
# Proof object simplifying inferences  : 34
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 23
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 38
# Processed clauses                    : 13250
# ...of these trivial                  : 45
# ...subsumed                          : 11001
# ...remaining for further processing  : 2204
# Other redundant clauses eliminated   : 10
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 312
# Backward-rewritten                   : 47
# Generated clauses                    : 106533
# ...of the previous two non-trivial   : 97613
# Contextual simplify-reflections      : 159
# Paramodulations                      : 106518
# Factorizations                       : 6
# Equation resolutions                 : 10
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1802
#    Positive orientable unit clauses  : 96
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 16
#    Non-unit-clauses                  : 1689
# Current number of unprocessed clauses: 82421
# ...number of literals in the above   : 325529
# Current number of archived formulas  : 0
# Current number of archived clauses   : 398
# Clause-clause subsumption calls (NU) : 705572
# Rec. Clause-clause subsumption calls : 367356
# Non-unit clause-clause subsumptions  : 7460
# Unit Clause-clause subsumption calls : 8996
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 775
# BW rewrite match successes           : 23
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1577949

# -------------------------------------------------
# User time                : 1.282 s
# System time              : 0.027 s
# Total time               : 1.308 s
# Maximum resident set size: 63396 pages
