# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t69_xboole_1, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>~((r1_tarski(X2,X1)&r1_xboole_0(X2,X1)))), file('data/MPTP/proofs/65303d0af9f74416ac33fb432aef0f86.E_input', t69_xboole_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/65303d0af9f74416ac33fb432aef0f86.E_input', t3_subset)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/65303d0af9f74416ac33fb432aef0f86.E_input', d1_xboole_0)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/65303d0af9f74416ac33fb432aef0f86.E_input', t3_xboole_0)).
fof(t17_yellow19, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_waybel_9(X1,k3_yellow19(X1,k2_struct_0(X1),X2),X3)<=>r1_waybel_7(X1,X2,X3))))), file('data/MPTP/proofs/65303d0af9f74416ac33fb432aef0f86.E_input', t17_yellow19)).
fof(d2_yellow_1, axiom, ![X1]:k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1)), file('data/MPTP/proofs/65303d0af9f74416ac33fb432aef0f86.E_input', d2_yellow_1)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/65303d0af9f74416ac33fb432aef0f86.E_input', dt_l1_pre_topc)).
fof(rc7_pre_topc, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>?[X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))&(~(v1_xboole_0(X2))&v4_pre_topc(X2,X1)))), file('data/MPTP/proofs/65303d0af9f74416ac33fb432aef0f86.E_input', rc7_pre_topc)).
fof(dt_k3_yellow19, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(~(v1_xboole_0(X3))&(v2_waybel_0(X3,k3_yellow_1(X2))&(v13_waybel_0(X3,k3_yellow_1(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))))))=>(~(v2_struct_0(k3_yellow19(X1,X2,X3)))&(v6_waybel_0(k3_yellow19(X1,X2,X3),X1)&l1_waybel_0(k3_yellow19(X1,X2,X3),X1)))), file('data/MPTP/proofs/65303d0af9f74416ac33fb432aef0f86.E_input', dt_k3_yellow19)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/65303d0af9f74416ac33fb432aef0f86.E_input', d3_struct_0)).
fof(fc4_yellow19, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(~(v1_xboole_0(X3))&(v2_waybel_0(X3,k3_yellow_1(X2))&(v13_waybel_0(X3,k3_yellow_1(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))))))=>(~(v2_struct_0(k3_yellow19(X1,X2,X3)))&(v3_orders_2(k3_yellow19(X1,X2,X3))&(v4_orders_2(k3_yellow19(X1,X2,X3))&v6_waybel_0(k3_yellow19(X1,X2,X3),X1))))), file('data/MPTP/proofs/65303d0af9f74416ac33fb432aef0f86.E_input', fc4_yellow19)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/65303d0af9f74416ac33fb432aef0f86.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/65303d0af9f74416ac33fb432aef0f86.E_input', d4_subset_1)).
fof(fc5_yellow19, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(~(v1_xboole_0(X3))&(v1_subset_1(X3,u1_struct_0(k3_yellow_1(X2)))&(v2_waybel_0(X3,k3_yellow_1(X2))&(v13_waybel_0(X3,k3_yellow_1(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))))))))=>(~(v2_struct_0(k3_yellow19(X1,X2,X3)))&(v6_waybel_0(k3_yellow19(X1,X2,X3),X1)&v7_waybel_0(k3_yellow19(X1,X2,X3))))), file('data/MPTP/proofs/65303d0af9f74416ac33fb432aef0f86.E_input', fc5_yellow19)).
fof(t12_yellow19, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_waybel_9(X1,X2,X3)<=>r1_waybel_7(X1,k2_yellow19(X1,X2),X3))))), file('data/MPTP/proofs/65303d0af9f74416ac33fb432aef0f86.E_input', t12_yellow19)).
fof(t15_yellow19, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2)))), file('data/MPTP/proofs/65303d0af9f74416ac33fb432aef0f86.E_input', t15_yellow19)).
fof(c_0_16, plain, ![X4, X5]:(v1_xboole_0(X5)|(~r1_tarski(X5,X4)|~r1_xboole_0(X5,X4))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t69_xboole_1])])])).
fof(c_0_17, plain, ![X6, X7]:((~m1_subset_1(X6,k1_zfmisc_1(X7))|r1_tarski(X6,X7))&(~r1_tarski(X6,X7)|m1_subset_1(X6,k1_zfmisc_1(X7)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_18, plain, ![X8, X9, X10]:((~v1_xboole_0(X8)|~r2_hidden(X9,X8))&(r2_hidden(esk1_1(X10),X10)|v1_xboole_0(X10))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_19, plain, ![X12, X13, X15, X16, X17]:(((r2_hidden(esk2_2(X12,X13),X12)|r1_xboole_0(X12,X13))&(r2_hidden(esk2_2(X12,X13),X13)|r1_xboole_0(X12,X13)))&(~r2_hidden(X17,X15)|~r2_hidden(X17,X16)|~r1_xboole_0(X15,X16))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_20, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_waybel_9(X1,k3_yellow19(X1,k2_struct_0(X1),X2),X3)<=>r1_waybel_7(X1,X2,X3)))))), inference(assume_negation,[status(cth)],[t17_yellow19])).
cnf(c_0_21, plain, (v1_xboole_0(X1)|~r1_tarski(X1,X2)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (r2_hidden(esk2_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_25, negated_conjecture, ((~v2_struct_0(esk3_0)&(v2_pre_topc(esk3_0)&l1_pre_topc(esk3_0)))&((~v1_xboole_0(esk4_0)&(v1_subset_1(esk4_0,u1_struct_0(k3_yellow_1(k2_struct_0(esk3_0))))&(v2_waybel_0(esk4_0,k3_yellow_1(k2_struct_0(esk3_0)))&(v13_waybel_0(esk4_0,k3_yellow_1(k2_struct_0(esk3_0)))&m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk3_0)))))))))&(m1_subset_1(esk5_0,u1_struct_0(esk3_0))&((~r3_waybel_9(esk3_0,k3_yellow19(esk3_0,k2_struct_0(esk3_0),esk4_0),esk5_0)|~r1_waybel_7(esk3_0,esk4_0,esk5_0))&(r3_waybel_9(esk3_0,k3_yellow19(esk3_0,k2_struct_0(esk3_0),esk4_0),esk5_0)|r1_waybel_7(esk3_0,esk4_0,esk5_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_20])])])])).
fof(c_0_26, plain, ![X24]:k3_yellow_1(X24)=k3_lattice3(k1_lattice3(X24)), inference(variable_rename,[status(thm)],[d2_yellow_1])).
fof(c_0_27, plain, ![X27]:(~l1_pre_topc(X27)|l1_struct_0(X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_28, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_29, plain, (r1_xboole_0(X1,X2)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_30, plain, ![X28]:((m1_subset_1(esk6_1(X28),k1_zfmisc_1(u1_struct_0(X28)))|(v2_struct_0(X28)|(~v2_pre_topc(X28)|~l1_pre_topc(X28))))&((~v1_xboole_0(esk6_1(X28))|(v2_struct_0(X28)|(~v2_pre_topc(X28)|~l1_pre_topc(X28))))&(v4_pre_topc(esk6_1(X28),X28)|(v2_struct_0(X28)|(~v2_pre_topc(X28)|~l1_pre_topc(X28)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[rc7_pre_topc])])])])])).
fof(c_0_31, plain, ![X21, X22, X23]:((~v2_struct_0(k3_yellow19(X21,X22,X23))|(v2_struct_0(X21)|~l1_struct_0(X21)|(v1_xboole_0(X22)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))|(v1_xboole_0(X23)|(~v2_waybel_0(X23,k3_yellow_1(X22))|(~v13_waybel_0(X23,k3_yellow_1(X22))|~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X22))))))))))&((v6_waybel_0(k3_yellow19(X21,X22,X23),X21)|(v2_struct_0(X21)|~l1_struct_0(X21)|(v1_xboole_0(X22)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))|(v1_xboole_0(X23)|(~v2_waybel_0(X23,k3_yellow_1(X22))|(~v13_waybel_0(X23,k3_yellow_1(X22))|~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X22))))))))))&(l1_waybel_0(k3_yellow19(X21,X22,X23),X21)|(v2_struct_0(X21)|~l1_struct_0(X21)|(v1_xboole_0(X22)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))|(v1_xboole_0(X23)|(~v2_waybel_0(X23,k3_yellow_1(X22))|(~v13_waybel_0(X23,k3_yellow_1(X22))|~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X22)))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_yellow19])])])])).
cnf(c_0_32, negated_conjecture, (v13_waybel_0(esk4_0,k3_yellow_1(k2_struct_0(esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_34, plain, ![X35]:(~l1_struct_0(X35)|k2_struct_0(X35)=u1_struct_0(X35)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_35, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36, negated_conjecture, (l1_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_37, negated_conjecture, (v2_waybel_0(esk4_0,k3_yellow_1(k2_struct_0(esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_38, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk3_0)))))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_39, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_40, plain, (m1_subset_1(esk6_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_41, plain, (v2_struct_0(X1)|~v1_xboole_0(esk6_1(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_42, plain, ![X30, X31, X32]:((~v2_struct_0(k3_yellow19(X30,X31,X32))|(v2_struct_0(X30)|~l1_struct_0(X30)|(v1_xboole_0(X31)|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))|(v1_xboole_0(X32)|(~v2_waybel_0(X32,k3_yellow_1(X31))|(~v13_waybel_0(X32,k3_yellow_1(X31))|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X31))))))))))&((v3_orders_2(k3_yellow19(X30,X31,X32))|(v2_struct_0(X30)|~l1_struct_0(X30)|(v1_xboole_0(X31)|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))|(v1_xboole_0(X32)|(~v2_waybel_0(X32,k3_yellow_1(X31))|(~v13_waybel_0(X32,k3_yellow_1(X31))|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X31))))))))))&((v4_orders_2(k3_yellow19(X30,X31,X32))|(v2_struct_0(X30)|~l1_struct_0(X30)|(v1_xboole_0(X31)|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))|(v1_xboole_0(X32)|(~v2_waybel_0(X32,k3_yellow_1(X31))|(~v13_waybel_0(X32,k3_yellow_1(X31))|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X31))))))))))&(v6_waybel_0(k3_yellow19(X30,X31,X32),X30)|(v2_struct_0(X30)|~l1_struct_0(X30)|(v1_xboole_0(X31)|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))|(v1_xboole_0(X32)|(~v2_waybel_0(X32,k3_yellow_1(X31))|(~v13_waybel_0(X32,k3_yellow_1(X31))|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X31))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc4_yellow19])])])])).
cnf(c_0_43, plain, (l1_waybel_0(k3_yellow19(X1,X2,X3),X1)|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_44, negated_conjecture, (v13_waybel_0(esk4_0,k3_lattice3(k1_lattice3(k2_struct_0(esk3_0))))), inference(rw,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_45, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_46, negated_conjecture, (l1_struct_0(esk3_0)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_47, negated_conjecture, (v2_waybel_0(esk4_0,k3_lattice3(k1_lattice3(k2_struct_0(esk3_0))))), inference(rw,[status(thm)],[c_0_37, c_0_33])).
cnf(c_0_48, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(esk3_0))))))), inference(rw,[status(thm)],[c_0_38, c_0_33])).
cnf(c_0_49, plain, (v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])).
cnf(c_0_50, negated_conjecture, (v2_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_51, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_52, plain, ![X25]:m1_subset_1(k2_subset_1(X25),k1_zfmisc_1(X25)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_53, plain, ![X26]:k2_subset_1(X26)=X26, inference(variable_rename,[status(thm)],[d4_subset_1])).
cnf(c_0_54, plain, (v4_orders_2(k3_yellow19(X1,X2,X3))|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_55, plain, ![X39, X40, X41]:((~v2_struct_0(k3_yellow19(X39,X40,X41))|(v2_struct_0(X39)|~l1_struct_0(X39)|(v1_xboole_0(X40)|~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X39)))|(v1_xboole_0(X41)|(~v1_subset_1(X41,u1_struct_0(k3_yellow_1(X40)))|(~v2_waybel_0(X41,k3_yellow_1(X40))|(~v13_waybel_0(X41,k3_yellow_1(X40))|~m1_subset_1(X41,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X40)))))))))))&((v6_waybel_0(k3_yellow19(X39,X40,X41),X39)|(v2_struct_0(X39)|~l1_struct_0(X39)|(v1_xboole_0(X40)|~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X39)))|(v1_xboole_0(X41)|(~v1_subset_1(X41,u1_struct_0(k3_yellow_1(X40)))|(~v2_waybel_0(X41,k3_yellow_1(X40))|(~v13_waybel_0(X41,k3_yellow_1(X40))|~m1_subset_1(X41,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X40)))))))))))&(v7_waybel_0(k3_yellow19(X39,X40,X41))|(v2_struct_0(X39)|~l1_struct_0(X39)|(v1_xboole_0(X40)|~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X39)))|(v1_xboole_0(X41)|(~v1_subset_1(X41,u1_struct_0(k3_yellow_1(X40)))|(~v2_waybel_0(X41,k3_yellow_1(X40))|(~v13_waybel_0(X41,k3_yellow_1(X40))|~m1_subset_1(X41,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X40))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc5_yellow19])])])])).
cnf(c_0_56, negated_conjecture, (v1_subset_1(esk4_0,u1_struct_0(k3_yellow_1(k2_struct_0(esk3_0))))), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_57, plain, ![X36, X37, X38]:((~r3_waybel_9(X36,X37,X38)|r1_waybel_7(X36,k2_yellow19(X36,X37),X38)|~m1_subset_1(X38,u1_struct_0(X36))|(v2_struct_0(X37)|(~v4_orders_2(X37)|(~v7_waybel_0(X37)|~l1_waybel_0(X37,X36))))|(v2_struct_0(X36)|(~v2_pre_topc(X36)|~l1_pre_topc(X36))))&(~r1_waybel_7(X36,k2_yellow19(X36,X37),X38)|r3_waybel_9(X36,X37,X38)|~m1_subset_1(X38,u1_struct_0(X36))|(v2_struct_0(X37)|(~v4_orders_2(X37)|(~v7_waybel_0(X37)|~l1_waybel_0(X37,X36))))|(v2_struct_0(X36)|(~v2_pre_topc(X36)|~l1_pre_topc(X36))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t12_yellow19])])])])])).
cnf(c_0_58, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|v2_struct_0(X1)|l1_waybel_0(k3_yellow19(X1,X2,X3),X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_33]), c_0_33]), c_0_33])).
cnf(c_0_59, negated_conjecture, (v13_waybel_0(esk4_0,k3_lattice3(k1_lattice3(u1_struct_0(esk3_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])])).
cnf(c_0_60, negated_conjecture, (v2_waybel_0(esk4_0,k3_lattice3(k1_lattice3(u1_struct_0(esk3_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_45]), c_0_46])])).
cnf(c_0_61, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(esk3_0))))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_45]), c_0_46])])).
cnf(c_0_62, negated_conjecture, (~v1_xboole_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_63, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_36])]), c_0_51])).
cnf(c_0_64, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_65, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_66, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|v2_struct_0(X1)|v4_orders_2(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_33]), c_0_33]), c_0_33])).
cnf(c_0_67, plain, (v7_waybel_0(k3_yellow19(X1,X2,X3))|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_subset_1(X3,u1_struct_0(k3_yellow_1(X2)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_68, negated_conjecture, (v1_subset_1(esk4_0,u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(esk3_0)))))), inference(rw,[status(thm)],[c_0_56, c_0_33])).
fof(c_0_69, plain, ![X33, X34]:(v2_struct_0(X33)|~l1_struct_0(X33)|(v1_xboole_0(X34)|(~v1_subset_1(X34,u1_struct_0(k3_yellow_1(k2_struct_0(X33))))|(~v2_waybel_0(X34,k3_yellow_1(k2_struct_0(X33)))|(~v13_waybel_0(X34,k3_yellow_1(k2_struct_0(X33)))|~m1_subset_1(X34,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X33))))))))|X34=k2_yellow19(X33,k3_yellow19(X33,k2_struct_0(X33),X34)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t15_yellow19])])])])).
cnf(c_0_70, plain, (r3_waybel_9(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_waybel_7(X1,k2_yellow19(X1,X2),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_71, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_72, negated_conjecture, (l1_waybel_0(k3_yellow19(X1,u1_struct_0(esk3_0),esk4_0),X1)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(X1)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_60]), c_0_61])]), c_0_62]), c_0_63])).
cnf(c_0_73, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_74, negated_conjecture, (v4_orders_2(k3_yellow19(X1,u1_struct_0(esk3_0),esk4_0))|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(X1)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_59]), c_0_60]), c_0_61])]), c_0_62]), c_0_63])).
cnf(c_0_75, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|v2_struct_0(X1)|v7_waybel_0(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v1_subset_1(X3,u1_struct_0(k3_lattice3(k1_lattice3(X2))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_33]), c_0_33]), c_0_33]), c_0_33])).
cnf(c_0_76, negated_conjecture, (v1_subset_1(esk4_0,u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(esk3_0)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_45]), c_0_46])])).
cnf(c_0_77, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2))|~l1_struct_0(X1)|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_78, negated_conjecture, (r3_waybel_9(esk3_0,X1,esk5_0)|v2_struct_0(X1)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~l1_waybel_0(X1,esk3_0)|~r1_waybel_7(esk3_0,k2_yellow19(esk3_0,X1),esk5_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_36]), c_0_50])]), c_0_51])).
cnf(c_0_79, negated_conjecture, (l1_waybel_0(k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0),esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_46])]), c_0_51])).
cnf(c_0_80, negated_conjecture, (v4_orders_2(k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_73]), c_0_46])]), c_0_51])).
cnf(c_0_81, negated_conjecture, (v7_waybel_0(k3_yellow19(X1,u1_struct_0(esk3_0),esk4_0))|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(X1)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_59]), c_0_60]), c_0_76]), c_0_61])]), c_0_62]), c_0_63])).
cnf(c_0_82, plain, (X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2))|v1_xboole_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~v1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(X1)))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(X1))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_33]), c_0_33]), c_0_33]), c_0_33])).
cnf(c_0_83, negated_conjecture, (r3_waybel_9(esk3_0,k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0),esk5_0)|v2_struct_0(k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0))|~v7_waybel_0(k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0))|~r1_waybel_7(esk3_0,k2_yellow19(esk3_0,k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0)),esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80])])).
cnf(c_0_84, negated_conjecture, (v7_waybel_0(k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_73]), c_0_46])]), c_0_51])).
cnf(c_0_85, negated_conjecture, (k2_yellow19(esk3_0,k3_yellow19(esk3_0,k2_struct_0(esk3_0),esk4_0))=esk4_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_44]), c_0_46]), c_0_47]), c_0_68]), c_0_48])]), c_0_51]), c_0_62])).
cnf(c_0_86, negated_conjecture, (r3_waybel_9(esk3_0,k3_yellow19(esk3_0,k2_struct_0(esk3_0),esk4_0),esk5_0)|r1_waybel_7(esk3_0,esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_87, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~v2_struct_0(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_88, negated_conjecture, (~r3_waybel_9(esk3_0,k3_yellow19(esk3_0,k2_struct_0(esk3_0),esk4_0),esk5_0)|~r1_waybel_7(esk3_0,esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_89, negated_conjecture, (r3_waybel_9(esk3_0,k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0),esk5_0)|v2_struct_0(k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0))|~r1_waybel_7(esk3_0,k2_yellow19(esk3_0,k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0)),esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_83, c_0_84])])).
cnf(c_0_90, negated_conjecture, (k2_yellow19(esk3_0,k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0))=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_45]), c_0_46])])).
cnf(c_0_91, negated_conjecture, (r1_waybel_7(esk3_0,esk4_0,esk5_0)|r3_waybel_9(esk3_0,k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0),esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_45]), c_0_46])])).
cnf(c_0_92, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~v2_struct_0(k3_yellow19(X1,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_33]), c_0_33]), c_0_33])).
cnf(c_0_93, negated_conjecture, (~r1_waybel_7(esk3_0,esk4_0,esk5_0)|~r3_waybel_9(esk3_0,k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0),esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_45]), c_0_46])])).
cnf(c_0_94, negated_conjecture, (r3_waybel_9(esk3_0,k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0),esk5_0)|v2_struct_0(k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0))), inference(csr,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_90]), c_0_91])).
cnf(c_0_95, plain, (r1_waybel_7(X1,k2_yellow19(X1,X2),X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_96, negated_conjecture, (v2_struct_0(X1)|~l1_struct_0(X1)|~v2_struct_0(k3_yellow19(X1,u1_struct_0(esk3_0),esk4_0))|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(X1)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_59]), c_0_60]), c_0_61])]), c_0_62]), c_0_63])).
cnf(c_0_97, negated_conjecture, (v2_struct_0(k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0))|~r1_waybel_7(esk3_0,esk4_0,esk5_0)), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_98, negated_conjecture, (r1_waybel_7(esk3_0,k2_yellow19(esk3_0,X1),esk5_0)|v2_struct_0(X1)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~l1_waybel_0(X1,esk3_0)|~r3_waybel_9(esk3_0,X1,esk5_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_71]), c_0_36]), c_0_50])]), c_0_51])).
cnf(c_0_99, negated_conjecture, (~r1_waybel_7(esk3_0,esk4_0,esk5_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_46]), c_0_73])]), c_0_51])).
cnf(c_0_100, negated_conjecture, (r1_waybel_7(esk3_0,k2_yellow19(esk3_0,k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0)),esk5_0)|v2_struct_0(k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0))|~v7_waybel_0(k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0))|~r3_waybel_9(esk3_0,k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0),esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_79]), c_0_80])])).
cnf(c_0_101, negated_conjecture, (r3_waybel_9(esk3_0,k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0),esk5_0)), inference(sr,[status(thm)],[c_0_91, c_0_99])).
cnf(c_0_102, negated_conjecture, (v2_struct_0(k3_yellow19(esk3_0,u1_struct_0(esk3_0),esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_100, c_0_84])]), c_0_90]), c_0_101])]), c_0_99])).
cnf(c_0_103, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_102]), c_0_46]), c_0_73])]), c_0_51]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 104
# Proof object clause steps            : 71
# Proof object formula steps           : 33
# Proof object conjectures             : 46
# Proof object clause conjectures      : 43
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 29
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 28
# Proof object simplifying inferences  : 108
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 38
# Processed clauses                    : 138
# ...of these trivial                  : 0
# ...subsumed                          : 18
# ...remaining for further processing  : 120
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 5
# Generated clauses                    : 74
# ...of the previous two non-trivial   : 69
# Contextual simplify-reflections      : 2
# Paramodulations                      : 72
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 77
#    Positive orientable unit clauses  : 23
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 48
# Current number of unprocessed clauses: 3
# ...number of literals in the above   : 10
# Current number of archived formulas  : 0
# Current number of archived clauses   : 45
# Clause-clause subsumption calls (NU) : 4852
# Rec. Clause-clause subsumption calls : 483
# Non-unit clause-clause subsumptions  : 18
# Unit Clause-clause subsumption calls : 36
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6126

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.007 s
# Total time               : 0.013 s
# Maximum resident set size: 28728 pages
