# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t39_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2))), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', t39_yellow_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', t24_orders_2)).
fof(t19_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))))), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', t19_yellow_0)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', redefinition_k6_domain_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', t75_enumset1)).
fof(t15_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4))))))), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', t15_yellow_0)).
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', d9_yellow_0)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', dt_k1_yellow_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', dt_l1_orders_2)).
fof(t38_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))&r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))))), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', t38_yellow_0)).
fof(t7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2)))))))), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', t7_yellow_0)).
fof(t16_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3))))))), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', t16_yellow_0)).
fof(d10_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_yellow_0(X1,X2)=>(X3=k2_yellow_0(X1,X2)<=>(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))))))), file('data/MPTP/proofs/6626b7eb9b6942f1b7c173bc142d023f.E_input', d10_yellow_0)).
fof(c_0_20, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))), introduced(definition)).
fof(c_0_21, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))), inference(split_equiv,[status(thm)],[c_0_20])).
fof(c_0_22, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2)))), inference(assume_negation,[status(cth)],[t39_yellow_0])).
fof(c_0_23, plain, ![X78, X79, X80, X81, X82]:(((r1_orders_2(X80,X81,X79)|(X81!=k11_lattice3(X80,X79,X78)|~r2_yellow_0(X80,k2_tarski(X79,X78)))|~m1_subset_1(X81,u1_struct_0(X80))|~epred1_3(X80,X79,X78))&((r1_orders_2(X80,X81,X78)|(X81!=k11_lattice3(X80,X79,X78)|~r2_yellow_0(X80,k2_tarski(X79,X78)))|~m1_subset_1(X81,u1_struct_0(X80))|~epred1_3(X80,X79,X78))&(~m1_subset_1(X82,u1_struct_0(X80))|(~r1_orders_2(X80,X82,X79)|~r1_orders_2(X80,X82,X78)|r1_orders_2(X80,X82,X81))|(X81!=k11_lattice3(X80,X79,X78)|~r2_yellow_0(X80,k2_tarski(X79,X78)))|~m1_subset_1(X81,u1_struct_0(X80))|~epred1_3(X80,X79,X78))))&(((X81=k11_lattice3(X80,X79,X78)|(m1_subset_1(esk9_4(X78,X79,X80,X81),u1_struct_0(X80))|~r1_orders_2(X80,X81,X78)|~r1_orders_2(X80,X81,X79))|~m1_subset_1(X81,u1_struct_0(X80))|~epred1_3(X80,X79,X78))&(r2_yellow_0(X80,k2_tarski(X79,X78))|(m1_subset_1(esk9_4(X78,X79,X80,X81),u1_struct_0(X80))|~r1_orders_2(X80,X81,X78)|~r1_orders_2(X80,X81,X79))|~m1_subset_1(X81,u1_struct_0(X80))|~epred1_3(X80,X79,X78)))&((((X81=k11_lattice3(X80,X79,X78)|(r1_orders_2(X80,esk9_4(X78,X79,X80,X81),X79)|~r1_orders_2(X80,X81,X78)|~r1_orders_2(X80,X81,X79))|~m1_subset_1(X81,u1_struct_0(X80))|~epred1_3(X80,X79,X78))&(r2_yellow_0(X80,k2_tarski(X79,X78))|(r1_orders_2(X80,esk9_4(X78,X79,X80,X81),X79)|~r1_orders_2(X80,X81,X78)|~r1_orders_2(X80,X81,X79))|~m1_subset_1(X81,u1_struct_0(X80))|~epred1_3(X80,X79,X78)))&((X81=k11_lattice3(X80,X79,X78)|(r1_orders_2(X80,esk9_4(X78,X79,X80,X81),X78)|~r1_orders_2(X80,X81,X78)|~r1_orders_2(X80,X81,X79))|~m1_subset_1(X81,u1_struct_0(X80))|~epred1_3(X80,X79,X78))&(r2_yellow_0(X80,k2_tarski(X79,X78))|(r1_orders_2(X80,esk9_4(X78,X79,X80,X81),X78)|~r1_orders_2(X80,X81,X78)|~r1_orders_2(X80,X81,X79))|~m1_subset_1(X81,u1_struct_0(X80))|~epred1_3(X80,X79,X78))))&((X81=k11_lattice3(X80,X79,X78)|(~r1_orders_2(X80,esk9_4(X78,X79,X80,X81),X81)|~r1_orders_2(X80,X81,X78)|~r1_orders_2(X80,X81,X79))|~m1_subset_1(X81,u1_struct_0(X80))|~epred1_3(X80,X79,X78))&(r2_yellow_0(X80,k2_tarski(X79,X78))|(~r1_orders_2(X80,esk9_4(X78,X79,X80,X81),X81)|~r1_orders_2(X80,X81,X78)|~r1_orders_2(X80,X81,X79))|~m1_subset_1(X81,u1_struct_0(X80))|~epred1_3(X80,X79,X78)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])])])).
fof(c_0_24, plain, ![X59, X60]:(v2_struct_0(X59)|(~v3_orders_2(X59)|~l1_orders_2(X59))|(~m1_subset_1(X60,u1_struct_0(X59))|r1_orders_2(X59,X60,X60))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_25, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_22])])])])).
cnf(c_0_26, plain, (X1=k11_lattice3(X2,X3,X4)|~r1_orders_2(X2,esk9_4(X4,X3,X2,X1),X1)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27, plain, (X1=k11_lattice3(X2,X3,X4)|r1_orders_2(X2,esk9_4(X4,X3,X2,X1),X4)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, plain, (X1=k11_lattice3(X2,X3,X4)|r1_orders_2(X2,esk9_4(X4,X3,X2,X1),X3)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_33, plain, (k11_lattice3(X1,X2,X3)=X3|~epred1_3(X1,X2,X3)|~r1_orders_2(X1,X3,X3)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_34, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])]), c_0_31])).
cnf(c_0_35, plain, (k11_lattice3(X1,X2,X3)=X2|~epred1_3(X1,X2,X3)|~r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_26, c_0_32])).
fof(c_0_36, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>epred1_3(X1,X2,X3)))), inference(apply_def,[status(thm)],[t19_yellow_0, c_0_20])).
fof(c_0_37, plain, ![X61, X62]:(v1_xboole_0(X61)|~m1_subset_1(X62,X61)|k6_domain_1(X61,X62)=k1_tarski(X62)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_38, plain, ![X32]:k2_tarski(X32,X32)=k1_tarski(X32), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_39, plain, ![X70, X71]:k1_enumset1(X70,X70,X71)=k2_tarski(X70,X71), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_40, plain, ![X22, X23, X24]:k2_enumset1(X22,X22,X23,X24)=k1_enumset1(X22,X23,X24), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_41, plain, ![X28, X29, X30, X31]:k3_enumset1(X28,X28,X29,X30,X31)=k2_enumset1(X28,X29,X30,X31), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_42, plain, ![X40, X41, X42, X43, X44]:k4_enumset1(X40,X40,X41,X42,X43,X44)=k3_enumset1(X40,X41,X42,X43,X44), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_43, plain, ![X72, X73, X74, X75, X76, X77]:k5_enumset1(X72,X72,X73,X74,X75,X76,X77)=k4_enumset1(X72,X73,X74,X75,X76,X77), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_44, plain, ![X15, X16, X17, X18, X19, X20, X21]:k6_enumset1(X15,X15,X16,X17,X18,X19,X20,X21)=k5_enumset1(X15,X16,X17,X18,X19,X20,X21), inference(variable_rename,[status(thm)],[t75_enumset1])).
cnf(c_0_45, negated_conjecture, (k11_lattice3(esk1_0,X1,X2)=X2|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_46, negated_conjecture, (k11_lattice3(esk1_0,X1,X2)=X1|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_35, c_0_34])).
fof(c_0_47, plain, ![X56, X57, X58]:(~v5_orders_2(X56)|~l1_orders_2(X56)|(~m1_subset_1(X57,u1_struct_0(X56))|(~m1_subset_1(X58,u1_struct_0(X56))|epred1_3(X56,X57,X58)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_36])])])).
cnf(c_0_48, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_49, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_52, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_53, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_54, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_55, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_56, negated_conjecture, (X1=X2|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,X1)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_57, plain, (epred1_3(X1,X2,X3)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_58, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_59, plain, ![X63, X64, X66, X67, X68]:(((m1_subset_1(esk7_2(X63,X64),u1_struct_0(X63))|~r1_yellow_0(X63,X64)|(~v5_orders_2(X63)|~l1_orders_2(X63)))&((r2_lattice3(X63,X64,esk7_2(X63,X64))|~r1_yellow_0(X63,X64)|(~v5_orders_2(X63)|~l1_orders_2(X63)))&(~m1_subset_1(X66,u1_struct_0(X63))|(~r2_lattice3(X63,X64,X66)|r1_orders_2(X63,esk7_2(X63,X64),X66))|~r1_yellow_0(X63,X64)|(~v5_orders_2(X63)|~l1_orders_2(X63)))))&((m1_subset_1(esk8_3(X63,X67,X68),u1_struct_0(X63))|~r2_lattice3(X63,X67,X68)|~m1_subset_1(X68,u1_struct_0(X63))|r1_yellow_0(X63,X67)|(~v5_orders_2(X63)|~l1_orders_2(X63)))&((r2_lattice3(X63,X67,esk8_3(X63,X67,X68))|~r2_lattice3(X63,X67,X68)|~m1_subset_1(X68,u1_struct_0(X63))|r1_yellow_0(X63,X67)|(~v5_orders_2(X63)|~l1_orders_2(X63)))&(~r1_orders_2(X63,X68,esk8_3(X63,X67,X68))|~r2_lattice3(X63,X67,X68)|~m1_subset_1(X68,u1_struct_0(X63))|r1_yellow_0(X63,X67)|(~v5_orders_2(X63)|~l1_orders_2(X63)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_yellow_0])])])])])])).
fof(c_0_60, plain, ![X35, X36, X37, X38]:(((r2_lattice3(X35,X36,X37)|X37!=k1_yellow_0(X35,X36)|~r1_yellow_0(X35,X36)|~m1_subset_1(X37,u1_struct_0(X35))|~l1_orders_2(X35))&(~m1_subset_1(X38,u1_struct_0(X35))|(~r2_lattice3(X35,X36,X38)|r1_orders_2(X35,X37,X38))|X37!=k1_yellow_0(X35,X36)|~r1_yellow_0(X35,X36)|~m1_subset_1(X37,u1_struct_0(X35))|~l1_orders_2(X35)))&((m1_subset_1(esk4_3(X35,X36,X37),u1_struct_0(X35))|~r2_lattice3(X35,X36,X37)|X37=k1_yellow_0(X35,X36)|~r1_yellow_0(X35,X36)|~m1_subset_1(X37,u1_struct_0(X35))|~l1_orders_2(X35))&((r2_lattice3(X35,X36,esk4_3(X35,X36,X37))|~r2_lattice3(X35,X36,X37)|X37=k1_yellow_0(X35,X36)|~r1_yellow_0(X35,X36)|~m1_subset_1(X37,u1_struct_0(X35))|~l1_orders_2(X35))&(~r1_orders_2(X35,X37,esk4_3(X35,X36,X37))|~r2_lattice3(X35,X36,X37)|X37=k1_yellow_0(X35,X36)|~r1_yellow_0(X35,X36)|~m1_subset_1(X37,u1_struct_0(X35))|~l1_orders_2(X35))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
fof(c_0_61, plain, ![X33, X34]:(~l1_orders_2(X33)|m1_subset_1(k1_yellow_0(X33,X34),u1_struct_0(X33))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
fof(c_0_62, plain, ![X55]:(v2_struct_0(X55)|~l1_struct_0(X55)|~v1_xboole_0(u1_struct_0(X55))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_63, plain, (k6_domain_1(X1,X2)=k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_49]), c_0_50]), c_0_51]), c_0_52]), c_0_53]), c_0_54]), c_0_55])).
cnf(c_0_64, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_65, negated_conjecture, (X1=X2|~r1_orders_2(esk1_0,X2,X1)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_30]), c_0_58])])).
cnf(c_0_66, plain, (r1_orders_2(X2,esk7_2(X2,X3),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r1_yellow_0(X2,X3)|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_67, plain, (r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~r1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_68, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_69, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_70, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_63, c_0_64])).
fof(c_0_71, plain, ![X54]:(~l1_orders_2(X54)|l1_struct_0(X54)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_72, negated_conjecture, (X1=esk7_2(esk1_0,X2)|~r1_yellow_0(esk1_0,X2)|~r2_lattice3(esk1_0,X2,X1)|~r1_orders_2(esk1_0,X1,esk7_2(esk1_0,X2))|~m1_subset_1(esk7_2(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_30]), c_0_58])])).
cnf(c_0_73, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),X3)|~r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_67]), c_0_68])).
cnf(c_0_74, plain, (r1_orders_2(X2,X1,X5)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,X1,X3)|~r1_orders_2(X2,X1,X4)|X5!=k11_lattice3(X2,X3,X4)|~r2_yellow_0(X2,k2_tarski(X3,X4))|~m1_subset_1(X5,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_75, plain, ![X45, X46]:((r1_yellow_0(X45,k6_domain_1(u1_struct_0(X45),X46))|~m1_subset_1(X46,u1_struct_0(X45))|(v2_struct_0(X45)|(~v3_orders_2(X45)|(~v5_orders_2(X45)|~l1_orders_2(X45)))))&(r2_yellow_0(X45,k6_domain_1(u1_struct_0(X45),X46))|~m1_subset_1(X46,u1_struct_0(X45))|(v2_struct_0(X45)|(~v3_orders_2(X45)|(~v5_orders_2(X45)|~l1_orders_2(X45)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t38_yellow_0])])])])])).
cnf(c_0_76, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_31])).
cnf(c_0_77, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_78, negated_conjecture, (k1_yellow_0(esk1_0,X1)=esk7_2(esk1_0,X2)|~r1_yellow_0(esk1_0,X2)|~r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,X2,k1_yellow_0(esk1_0,X1))|~r2_lattice3(esk1_0,X1,esk7_2(esk1_0,X2))|~m1_subset_1(esk7_2(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_30])])).
cnf(c_0_79, plain, (r2_lattice3(X1,X2,esk7_2(X1,X2))|~r1_yellow_0(X1,X2)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_80, plain, (r2_lattice3(X1,X2,X3)|X3!=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_81, plain, (r1_orders_2(X2,X1,X5)|X5!=k11_lattice3(X2,X3,X4)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~epred1_3(X2,X3,X4)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~r2_yellow_0(X2,k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_50]), c_0_51]), c_0_52]), c_0_53]), c_0_54]), c_0_55])).
cnf(c_0_82, plain, (r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_83, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_30])])).
cnf(c_0_84, negated_conjecture, (esk7_2(esk1_0,X1)=k1_yellow_0(esk1_0,X1)|~r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,X1))|~m1_subset_1(esk7_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_30]), c_0_58])])).
cnf(c_0_85, plain, (r2_lattice3(X1,X2,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_80]), c_0_68])).
cnf(c_0_86, plain, (r1_orders_2(X1,X2,k11_lattice3(X1,X3,X4))|~epred1_3(X1,X3,X4)|~r1_orders_2(X1,X2,X4)|~r1_orders_2(X1,X2,X3)|~r2_yellow_0(X1,k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X4))|~m1_subset_1(k11_lattice3(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(er,[status(thm)],[c_0_81])).
cnf(c_0_87, negated_conjecture, (r2_yellow_0(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_64]), c_0_30]), c_0_58]), c_0_29])]), c_0_31])).
fof(c_0_88, plain, ![X25, X26, X27]:((~r1_lattice3(X25,k1_tarski(X27),X26)|r1_orders_2(X25,X26,X27)|~m1_subset_1(X27,u1_struct_0(X25))|~m1_subset_1(X26,u1_struct_0(X25))|~l1_orders_2(X25))&((~r1_orders_2(X25,X26,X27)|r1_lattice3(X25,k1_tarski(X27),X26)|~m1_subset_1(X27,u1_struct_0(X25))|~m1_subset_1(X26,u1_struct_0(X25))|~l1_orders_2(X25))&((~r2_lattice3(X25,k1_tarski(X27),X26)|r1_orders_2(X25,X27,X26)|~m1_subset_1(X27,u1_struct_0(X25))|~m1_subset_1(X26,u1_struct_0(X25))|~l1_orders_2(X25))&(~r1_orders_2(X25,X27,X26)|r2_lattice3(X25,k1_tarski(X27),X26)|~m1_subset_1(X27,u1_struct_0(X25))|~m1_subset_1(X26,u1_struct_0(X25))|~l1_orders_2(X25))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])).
cnf(c_0_89, negated_conjecture, (esk7_2(esk1_0,X1)=k1_yellow_0(esk1_0,X1)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk7_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_30])])).
cnf(c_0_90, plain, (m1_subset_1(esk7_2(X1,X2),u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_91, negated_conjecture, (r1_orders_2(esk1_0,X1,k11_lattice3(esk1_0,esk2_0,esk2_0))|~epred1_3(esk1_0,esk2_0,esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(k11_lattice3(esk1_0,esk2_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_92, plain, (r1_orders_2(X1,X2,X3)|~r2_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_93, negated_conjecture, (esk7_2(esk1_0,X1)=k1_yellow_0(esk1_0,X1)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_30]), c_0_58])])).
cnf(c_0_94, negated_conjecture, (k11_lattice3(esk1_0,esk2_0,esk2_0)=X1|~epred1_3(esk1_0,esk2_0,esk2_0)|~r1_orders_2(esk1_0,k11_lattice3(esk1_0,esk2_0,esk2_0),X1)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(k11_lattice3(esk1_0,esk2_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_65, c_0_91])).
cnf(c_0_95, plain, (r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r2_lattice3(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2),X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_49]), c_0_50]), c_0_51]), c_0_52]), c_0_53]), c_0_54]), c_0_55])).
cnf(c_0_96, negated_conjecture, (r2_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_93]), c_0_30]), c_0_58])])).
cnf(c_0_97, negated_conjecture, (esk2_0=X1|~epred1_3(esk1_0,esk2_0,esk2_0)|~r1_orders_2(esk1_0,esk2_0,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_45]), c_0_64])])).
cnf(c_0_98, negated_conjecture, (X1=k1_yellow_0(esk1_0,X2)|~r1_yellow_0(esk1_0,X2)|~r2_lattice3(esk1_0,X2,X1)|~r1_orders_2(esk1_0,X1,k1_yellow_0(esk1_0,X2))|~m1_subset_1(k1_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_73]), c_0_30])])).
cnf(c_0_99, negated_conjecture, (r1_orders_2(esk1_0,X1,k1_yellow_0(esk1_0,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)))|~r1_yellow_0(esk1_0,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))|~m1_subset_1(k1_yellow_0(esk1_0,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_30])])).
cnf(c_0_100, negated_conjecture, (esk2_0=X1|~r1_orders_2(esk1_0,esk2_0,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_57]), c_0_64]), c_0_30]), c_0_58])])).
cnf(c_0_101, plain, (r1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_102, negated_conjecture, (k1_yellow_0(esk1_0,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))=X1|~r1_yellow_0(esk1_0,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))|~r2_lattice3(esk1_0,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),X1)|~m1_subset_1(k1_yellow_0(esk1_0,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_103, plain, (r2_lattice3(X1,k1_tarski(X2),X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
fof(c_0_104, plain, ![X47, X48, X50, X51, X52]:(((m1_subset_1(esk5_2(X47,X48),u1_struct_0(X47))|~r2_yellow_0(X47,X48)|(~v5_orders_2(X47)|~l1_orders_2(X47)))&((r1_lattice3(X47,X48,esk5_2(X47,X48))|~r2_yellow_0(X47,X48)|(~v5_orders_2(X47)|~l1_orders_2(X47)))&(~m1_subset_1(X50,u1_struct_0(X47))|(~r1_lattice3(X47,X48,X50)|r1_orders_2(X47,X50,esk5_2(X47,X48)))|~r2_yellow_0(X47,X48)|(~v5_orders_2(X47)|~l1_orders_2(X47)))))&((m1_subset_1(esk6_3(X47,X51,X52),u1_struct_0(X47))|~r1_lattice3(X47,X51,X52)|~m1_subset_1(X52,u1_struct_0(X47))|r2_yellow_0(X47,X51)|(~v5_orders_2(X47)|~l1_orders_2(X47)))&((r1_lattice3(X47,X51,esk6_3(X47,X51,X52))|~r1_lattice3(X47,X51,X52)|~m1_subset_1(X52,u1_struct_0(X47))|r2_yellow_0(X47,X51)|(~v5_orders_2(X47)|~l1_orders_2(X47)))&(~r1_orders_2(X47,esk6_3(X47,X51,X52),X52)|~r1_lattice3(X47,X51,X52)|~m1_subset_1(X52,u1_struct_0(X47))|r2_yellow_0(X47,X51)|(~v5_orders_2(X47)|~l1_orders_2(X47)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_yellow_0])])])])])])).
cnf(c_0_105, plain, (r1_orders_2(X1,X3,X2)|~r1_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_106, negated_conjecture, (esk2_0=X1|~r1_orders_2(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_34]), c_0_64])])).
cnf(c_0_107, negated_conjecture, (r1_yellow_0(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_83]), c_0_64]), c_0_30]), c_0_58]), c_0_29])]), c_0_31])).
cnf(c_0_108, negated_conjecture, (k1_yellow_0(esk1_0,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))=X1|~r1_yellow_0(esk1_0,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))|~r2_lattice3(esk1_0,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_68]), c_0_30])])).
cnf(c_0_109, plain, (r2_lattice3(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2),X3)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_49]), c_0_50]), c_0_51]), c_0_52]), c_0_53]), c_0_54]), c_0_55])).
cnf(c_0_110, plain, (r1_orders_2(X2,X1,esk5_2(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~r2_yellow_0(X2,X3)|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_111, plain, (r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2),X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_105, c_0_49]), c_0_50]), c_0_51]), c_0_52]), c_0_53]), c_0_54]), c_0_55])).
cnf(c_0_112, plain, (r1_lattice3(X1,X2,esk5_2(X1,X2))|~r2_yellow_0(X1,X2)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_113, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_114, negated_conjecture, (k1_yellow_0(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))=esk2_0|~r1_orders_2(esk1_0,k1_yellow_0(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)),esk2_0)|~m1_subset_1(k1_yellow_0(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_99]), c_0_107]), c_0_64])])).
cnf(c_0_115, negated_conjecture, (k1_yellow_0(esk1_0,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))=X1|~r1_yellow_0(esk1_0,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_30])]), c_0_34])).
cnf(c_0_116, plain, (k6_enumset1(k1_yellow_0(X1,X2),k1_yellow_0(X1,X2),k1_yellow_0(X1,X2),k1_yellow_0(X1,X2),k1_yellow_0(X1,X2),k1_yellow_0(X1,X2),k1_yellow_0(X1,X2),k1_yellow_0(X1,X2))=k6_domain_1(u1_struct_0(X1),k1_yellow_0(X1,X2))|v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_63, c_0_68])).
fof(c_0_117, plain, ![X10, X11, X12, X13]:(((r1_lattice3(X10,X11,X12)|X12!=k2_yellow_0(X10,X11)|~r2_yellow_0(X10,X11)|~m1_subset_1(X12,u1_struct_0(X10))|~l1_orders_2(X10))&(~m1_subset_1(X13,u1_struct_0(X10))|(~r1_lattice3(X10,X11,X13)|r1_orders_2(X10,X13,X12))|X12!=k2_yellow_0(X10,X11)|~r2_yellow_0(X10,X11)|~m1_subset_1(X12,u1_struct_0(X10))|~l1_orders_2(X10)))&((m1_subset_1(esk3_3(X10,X11,X12),u1_struct_0(X10))|~r1_lattice3(X10,X11,X12)|X12=k2_yellow_0(X10,X11)|~r2_yellow_0(X10,X11)|~m1_subset_1(X12,u1_struct_0(X10))|~l1_orders_2(X10))&((r1_lattice3(X10,X11,esk3_3(X10,X11,X12))|~r1_lattice3(X10,X11,X12)|X12=k2_yellow_0(X10,X11)|~r2_yellow_0(X10,X11)|~m1_subset_1(X12,u1_struct_0(X10))|~l1_orders_2(X10))&(~r1_orders_2(X10,esk3_3(X10,X11,X12),X12)|~r1_lattice3(X10,X11,X12)|X12=k2_yellow_0(X10,X11)|~r2_yellow_0(X10,X11)|~m1_subset_1(X12,u1_struct_0(X10))|~l1_orders_2(X10))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_yellow_0])])])])])).
cnf(c_0_118, negated_conjecture, (esk5_2(esk1_0,X1)=esk2_0|~r1_orders_2(esk1_0,esk5_2(esk1_0,X1),esk2_0)|~r1_lattice3(esk1_0,X1,esk2_0)|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(esk5_2(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_110]), c_0_64]), c_0_30]), c_0_58])])).
cnf(c_0_119, plain, (r1_orders_2(X1,esk5_2(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2)),X2)|~r2_yellow_0(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_113])).
cnf(c_0_120, negated_conjecture, (k1_yellow_0(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))=esk2_0|~r2_lattice3(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),esk2_0)|~m1_subset_1(k1_yellow_0(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_73]), c_0_107]), c_0_64]), c_0_30])])).
cnf(c_0_121, negated_conjecture, (r2_lattice3(esk1_0,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),X1)|~r1_yellow_0(esk1_0,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_115]), c_0_30])])).
cnf(c_0_122, negated_conjecture, (k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(X1),k1_yellow_0(X1,X2)))=k1_yellow_0(X1,X2)|v1_xboole_0(u1_struct_0(X1))|~r1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(X1),k1_yellow_0(X1,X2)))|~m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(esk1_0))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_115, c_0_116])).
cnf(c_0_123, plain, (X3=k2_yellow_0(X1,X2)|~r1_orders_2(X1,esk3_3(X1,X2,X3),X3)|~r1_lattice3(X1,X2,X3)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_117])).
cnf(c_0_124, negated_conjecture, (esk5_2(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))=esk2_0|~r1_lattice3(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),esk2_0)|~m1_subset_1(esk5_2(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_87]), c_0_64]), c_0_30]), c_0_58])])).
cnf(c_0_125, negated_conjecture, (k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_126, negated_conjecture, (k1_yellow_0(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))=esk2_0|~r2_lattice3(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_68]), c_0_30])])).
cnf(c_0_127, negated_conjecture, (r2_lattice3(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_107]), c_0_64])])).
cnf(c_0_128, negated_conjecture, (k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1))=X1|v1_xboole_0(u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))|~r1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_115]), c_0_30])])).
cnf(c_0_129, plain, (esk5_2(X1,X2)=k2_yellow_0(X1,X3)|~r1_lattice3(X1,X2,esk3_3(X1,X3,esk5_2(X1,X2)))|~r1_lattice3(X1,X3,esk5_2(X1,X2))|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2)|~m1_subset_1(esk3_3(X1,X3,esk5_2(X1,X2)),u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_110]), c_0_113])).
cnf(c_0_130, negated_conjecture, (esk5_2(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))=esk2_0|~r1_lattice3(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_113]), c_0_87]), c_0_30]), c_0_58])])).
cnf(c_0_131, negated_conjecture, (k1_yellow_0(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))!=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_125, c_0_83]), c_0_83])).
cnf(c_0_132, negated_conjecture, (k1_yellow_0(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_126, c_0_127])])).
cnf(c_0_133, negated_conjecture, (k1_yellow_0(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))=esk2_0|v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_107]), c_0_83]), c_0_83]), c_0_107]), c_0_64])])).
cnf(c_0_134, negated_conjecture, (k2_yellow_0(esk1_0,X1)=esk2_0|~r1_lattice3(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),esk3_3(esk1_0,X1,esk2_0))|~r1_lattice3(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),esk2_0)|~r1_lattice3(esk1_0,X1,esk2_0)|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(esk3_3(esk1_0,X1,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130]), c_0_87]), c_0_30]), c_0_58])])).
cnf(c_0_135, plain, (r1_lattice3(X1,X2,esk3_3(X1,X2,X3))|X3=k2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_117])).
cnf(c_0_136, negated_conjecture, (k2_yellow_0(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_131, c_0_132])])).
cnf(c_0_137, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_133]), c_0_107]), c_0_64])])).
cnf(c_0_138, negated_conjecture, (~r1_lattice3(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),esk2_0)|~m1_subset_1(esk3_3(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_87]), c_0_64]), c_0_30])]), c_0_136])).
cnf(c_0_139, plain, (m1_subset_1(esk3_3(X1,X2,X3),u1_struct_0(X1))|X3=k2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_117])).
cnf(c_0_140, plain, (r1_lattice3(X1,k1_tarski(X3),X2)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_141, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_137]), c_0_31])).
cnf(c_0_142, negated_conjecture, (~r1_lattice3(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_87]), c_0_64]), c_0_30])]), c_0_136])).
cnf(c_0_143, plain, (r1_lattice3(X1,k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X3),X2)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_140, c_0_49]), c_0_50]), c_0_51]), c_0_52]), c_0_53]), c_0_54]), c_0_55])).
cnf(c_0_144, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_77]), c_0_30])])).
cnf(c_0_145, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_143]), c_0_144]), c_0_64]), c_0_30])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 146
# Proof object clause steps            : 101
# Proof object formula steps           : 45
# Proof object conjectures             : 56
# Proof object clause conjectures      : 53
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 40
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 49
# Proof object simplifying inferences  : 165
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 20
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 58
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 51
# Processed clauses                    : 405
# ...of these trivial                  : 2
# ...subsumed                          : 132
# ...remaining for further processing  : 271
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 26
# Backward-rewritten                   : 32
# Generated clauses                    : 686
# ...of the previous two non-trivial   : 609
# Contextual simplify-reflections      : 17
# Paramodulations                      : 679
# Factorizations                       : 0
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 155
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 141
# Current number of unprocessed clauses: 286
# ...number of literals in the above   : 2285
# Current number of archived formulas  : 0
# Current number of archived clauses   : 116
# Clause-clause subsumption calls (NU) : 3878
# Rec. Clause-clause subsumption calls : 551
# Non-unit clause-clause subsumptions  : 173
# Unit Clause-clause subsumption calls : 358
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 21
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 34473

# -------------------------------------------------
# User time                : 0.031 s
# System time              : 0.000 s
# Total time               : 0.031 s
# Maximum resident set size: 28716 pages
