# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>(r2_hidden(X1,k1_relat_1(X3))&X2=k1_funct_1(X3,X1)))), file('data/MPTP/proofs/66441ab5796b4182901fd4fad80a9ddd.E_input', t8_funct_1)).
fof(d8_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(X3=k5_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>?[X6]:(r2_hidden(k4_tarski(X4,X6),X1)&r2_hidden(k4_tarski(X6,X5),X2))))))), file('data/MPTP/proofs/66441ab5796b4182901fd4fad80a9ddd.E_input', d8_relat_1)).
fof(t143_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k9_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k1_relat_1(X3))&(r2_hidden(k4_tarski(X4,X1),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/66441ab5796b4182901fd4fad80a9ddd.E_input', t143_relat_1)).
fof(t44_funct_1, conjecture, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((k2_relat_1(X1)=k1_relat_1(X2)&k5_relat_1(X1,X2)=X1)=>X2=k6_relat_1(k1_relat_1(X2))))), file('data/MPTP/proofs/66441ab5796b4182901fd4fad80a9ddd.E_input', t44_funct_1)).
fof(t146_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k9_relat_1(X1,k1_relat_1(X1))=k2_relat_1(X1)), file('data/MPTP/proofs/66441ab5796b4182901fd4fad80a9ddd.E_input', t146_relat_1)).
fof(t34_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(X2=k6_relat_1(X1)<=>(k1_relat_1(X2)=X1&![X3]:(r2_hidden(X3,X1)=>k1_funct_1(X2,X3)=X3)))), file('data/MPTP/proofs/66441ab5796b4182901fd4fad80a9ddd.E_input', t34_funct_1)).
fof(c_0_6, plain, ![X13, X14, X15]:(((r2_hidden(X13,k1_relat_1(X15))|~r2_hidden(k4_tarski(X13,X14),X15)|(~v1_relat_1(X15)|~v1_funct_1(X15)))&(X14=k1_funct_1(X15,X13)|~r2_hidden(k4_tarski(X13,X14),X15)|(~v1_relat_1(X15)|~v1_funct_1(X15))))&(~r2_hidden(X13,k1_relat_1(X15))|X14!=k1_funct_1(X15,X13)|r2_hidden(k4_tarski(X13,X14),X15)|(~v1_relat_1(X15)|~v1_funct_1(X15)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_funct_1])])])).
fof(c_0_7, plain, ![X17, X18, X19, X20, X21, X23, X24, X25, X28]:((((r2_hidden(k4_tarski(X20,esk4_5(X17,X18,X19,X20,X21)),X17)|~r2_hidden(k4_tarski(X20,X21),X19)|X19!=k5_relat_1(X17,X18)|~v1_relat_1(X19)|~v1_relat_1(X18)|~v1_relat_1(X17))&(r2_hidden(k4_tarski(esk4_5(X17,X18,X19,X20,X21),X21),X18)|~r2_hidden(k4_tarski(X20,X21),X19)|X19!=k5_relat_1(X17,X18)|~v1_relat_1(X19)|~v1_relat_1(X18)|~v1_relat_1(X17)))&(~r2_hidden(k4_tarski(X23,X25),X17)|~r2_hidden(k4_tarski(X25,X24),X18)|r2_hidden(k4_tarski(X23,X24),X19)|X19!=k5_relat_1(X17,X18)|~v1_relat_1(X19)|~v1_relat_1(X18)|~v1_relat_1(X17)))&((~r2_hidden(k4_tarski(esk5_3(X17,X18,X19),esk6_3(X17,X18,X19)),X19)|(~r2_hidden(k4_tarski(esk5_3(X17,X18,X19),X28),X17)|~r2_hidden(k4_tarski(X28,esk6_3(X17,X18,X19)),X18))|X19=k5_relat_1(X17,X18)|~v1_relat_1(X19)|~v1_relat_1(X18)|~v1_relat_1(X17))&((r2_hidden(k4_tarski(esk5_3(X17,X18,X19),esk7_3(X17,X18,X19)),X17)|r2_hidden(k4_tarski(esk5_3(X17,X18,X19),esk6_3(X17,X18,X19)),X19)|X19=k5_relat_1(X17,X18)|~v1_relat_1(X19)|~v1_relat_1(X18)|~v1_relat_1(X17))&(r2_hidden(k4_tarski(esk7_3(X17,X18,X19),esk6_3(X17,X18,X19)),X18)|r2_hidden(k4_tarski(esk5_3(X17,X18,X19),esk6_3(X17,X18,X19)),X19)|X19=k5_relat_1(X17,X18)|~v1_relat_1(X19)|~v1_relat_1(X18)|~v1_relat_1(X17))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_relat_1])])])])])])).
cnf(c_0_8, plain, (X1=k1_funct_1(X2,X3)|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, plain, (r2_hidden(k4_tarski(esk4_5(X1,X2,X3,X4,X5),X5),X2)|~r2_hidden(k4_tarski(X4,X5),X3)|X3!=k5_relat_1(X1,X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10, plain, (r2_hidden(k4_tarski(X1,esk4_5(X2,X3,X4,X1,X5)),X2)|~r2_hidden(k4_tarski(X1,X5),X4)|X4!=k5_relat_1(X2,X3)|~v1_relat_1(X4)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (k1_funct_1(X1,esk4_5(X2,X1,X3,X4,X5))=X5|X3!=k5_relat_1(X2,X1)|~r2_hidden(k4_tarski(X4,X5),X3)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_12, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X1,X4)|X3!=k5_relat_1(X1,X2)|~r2_hidden(k4_tarski(X4,X5),X3)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_8, c_0_10])).
fof(c_0_13, plain, ![X30, X31, X32, X34]:(((r2_hidden(esk8_3(X30,X31,X32),k1_relat_1(X32))|~r2_hidden(X30,k9_relat_1(X32,X31))|~v1_relat_1(X32))&((r2_hidden(k4_tarski(esk8_3(X30,X31,X32),X30),X32)|~r2_hidden(X30,k9_relat_1(X32,X31))|~v1_relat_1(X32))&(r2_hidden(esk8_3(X30,X31,X32),X31)|~r2_hidden(X30,k9_relat_1(X32,X31))|~v1_relat_1(X32))))&(~r2_hidden(X34,k1_relat_1(X32))|(~r2_hidden(k4_tarski(X34,X30),X32)|~r2_hidden(X34,X31))|r2_hidden(X30,k9_relat_1(X32,X31))|~v1_relat_1(X32))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t143_relat_1])])])])])).
fof(c_0_14, negated_conjecture, ~(![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((k2_relat_1(X1)=k1_relat_1(X2)&k5_relat_1(X1,X2)=X1)=>X2=k6_relat_1(k1_relat_1(X2)))))), inference(assume_negation,[status(cth)],[t44_funct_1])).
cnf(c_0_15, plain, (k1_funct_1(X1,k1_funct_1(X2,X3))=X4|X5!=k5_relat_1(X2,X1)|~r2_hidden(k4_tarski(X3,X4),X5)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X5)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16, plain, (r2_hidden(k4_tarski(esk8_3(X1,X2,X3),X1),X3)|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_17, negated_conjecture, ((v1_relat_1(esk1_0)&v1_funct_1(esk1_0))&((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&((k2_relat_1(esk1_0)=k1_relat_1(esk2_0)&k5_relat_1(esk1_0,esk2_0)=esk1_0)&esk2_0!=k6_relat_1(k1_relat_1(esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
fof(c_0_18, plain, ![X16]:(~v1_relat_1(X16)|k9_relat_1(X16,k1_relat_1(X16))=k2_relat_1(X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t146_relat_1])])).
cnf(c_0_19, plain, (k1_funct_1(X1,k1_funct_1(X2,esk8_3(X3,X4,X5)))=X3|X5!=k5_relat_1(X2,X1)|~r2_hidden(X3,k9_relat_1(X5,X4))|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X5)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_20, plain, (k1_funct_1(X1,esk8_3(X2,X3,X1))=X2|~r2_hidden(X2,k9_relat_1(X1,X3))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_8, c_0_16])).
cnf(c_0_21, negated_conjecture, (k2_relat_1(esk1_0)=k1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, plain, (k9_relat_1(X1,k1_relat_1(X1))=k2_relat_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_24, plain, ![X9, X10, X11]:(((k1_relat_1(X10)=X9|X10!=k6_relat_1(X9)|(~v1_relat_1(X10)|~v1_funct_1(X10)))&(~r2_hidden(X11,X9)|k1_funct_1(X10,X11)=X11|X10!=k6_relat_1(X9)|(~v1_relat_1(X10)|~v1_funct_1(X10))))&((r2_hidden(esk3_2(X9,X10),X9)|k1_relat_1(X10)!=X9|X10=k6_relat_1(X9)|(~v1_relat_1(X10)|~v1_funct_1(X10)))&(k1_funct_1(X10,esk3_2(X9,X10))!=esk3_2(X9,X10)|k1_relat_1(X10)!=X9|X10=k6_relat_1(X9)|(~v1_relat_1(X10)|~v1_funct_1(X10))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t34_funct_1])])])])])).
cnf(c_0_25, plain, (k1_funct_1(X1,X2)=X2|k5_relat_1(X3,X1)!=X3|~r2_hidden(X2,k9_relat_1(X3,X4))|~v1_funct_1(X1)|~v1_funct_1(X3)|~v1_relat_1(X1)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_26, negated_conjecture, (k9_relat_1(esk1_0,k1_relat_1(esk1_0))=k1_relat_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])])).
cnf(c_0_27, negated_conjecture, (v1_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28, plain, (X1=k6_relat_1(X2)|k1_funct_1(X1,esk3_2(X2,X1))!=esk3_2(X2,X1)|k1_relat_1(X1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29, negated_conjecture, (k1_funct_1(X1,X2)=X2|k5_relat_1(esk1_0,X1)!=esk1_0|~r2_hidden(X2,k1_relat_1(esk2_0))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_23])])).
cnf(c_0_30, negated_conjecture, (X1=k6_relat_1(X2)|k5_relat_1(esk1_0,X1)!=esk1_0|k1_relat_1(X1)!=X2|~r2_hidden(esk3_2(X2,X1),k1_relat_1(esk2_0))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_31, plain, (r2_hidden(esk3_2(X1,X2),X1)|X2=k6_relat_1(X1)|k1_relat_1(X2)!=X1|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, negated_conjecture, (X1=k6_relat_1(k1_relat_1(esk2_0))|k5_relat_1(esk1_0,X1)!=esk1_0|k1_relat_1(X1)!=k1_relat_1(esk2_0)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_33, negated_conjecture, (k5_relat_1(esk1_0,esk2_0)=esk1_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_34, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_35, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_36, negated_conjecture, (esk2_0!=k6_relat_1(k1_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_37, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34]), c_0_35])]), c_0_36]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 38
# Proof object clause steps            : 25
# Proof object formula steps           : 13
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 9
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 25
# Processed clauses                    : 13814
# ...of these trivial                  : 1005
# ...subsumed                          : 10311
# ...remaining for further processing  : 2498
# Other redundant clauses eliminated   : 588
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 268
# Backward-rewritten                   : 0
# Generated clauses                    : 397027
# ...of the previous two non-trivial   : 383086
# Contextual simplify-reflections      : 77
# Paramodulations                      : 396310
# Factorizations                       : 0
# Equation resolutions                 : 717
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2205
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 2196
# Current number of unprocessed clauses: 368255
# ...number of literals in the above   : 4102890
# Current number of archived formulas  : 0
# Current number of archived clauses   : 293
# Clause-clause subsumption calls (NU) : 2337007
# Rec. Clause-clause subsumption calls : 86425
# Non-unit clause-clause subsumptions  : 9358
# Unit Clause-clause subsumption calls : 206
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 15498233

# -------------------------------------------------
# User time                : 7.801 s
# System time              : 0.116 s
# Total time               : 7.917 s
# Maximum resident set size: 370140 pages
