# No SInE strategy applied
# Auto-Mode selected heuristic H_____047_B31_F1_PI_AE_R4_CS_SP_S4c
# and selection function SelectCQPrecWNTNp.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t96_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,u1_struct_0(X2))=>X4=k1_funct_1(X3,X4)))=>r2_funct_2(u1_struct_0(X2),u1_struct_0(X1),k4_tmap_1(X1,X2),X3))))), file('data/MPTP/proofs/66df56672ed048759d4dbc07d81addf0.E_input', t96_tmap_1)).
fof(dt_k4_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_pre_topc(X2,X1))=>(v1_funct_1(k4_tmap_1(X1,X2))&(v1_funct_2(k4_tmap_1(X1,X2),u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(k4_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))), file('data/MPTP/proofs/66df56672ed048759d4dbc07d81addf0.E_input', dt_k4_tmap_1)).
fof(d9_funct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(r2_funct_2(X1,X2,X3,X4)<=>![X5]:(m1_subset_1(X5,X1)=>k1_funct_1(X3,X5)=k1_funct_1(X4,X5))))), file('data/MPTP/proofs/66df56672ed048759d4dbc07d81addf0.E_input', d9_funct_2)).
fof(t59_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X2))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X1))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))))=>(![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>(r2_hidden(X6,u1_struct_0(X3))=>k3_funct_2(u1_struct_0(X2),u1_struct_0(X1),X4,X6)=k1_funct_1(X5,X6)))=>r2_funct_2(u1_struct_0(X3),u1_struct_0(X1),k2_tmap_1(X2,X1,X4,X3),X5))))))), file('data/MPTP/proofs/66df56672ed048759d4dbc07d81addf0.E_input', t59_tmap_1)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/66df56672ed048759d4dbc07d81addf0.E_input', dt_m1_pre_topc)).
fof(cc1_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>v2_pre_topc(X2))), file('data/MPTP/proofs/66df56672ed048759d4dbc07d81addf0.E_input', cc1_pre_topc)).
fof(dt_k3_tmap_1, axiom, ![X1, X2, X3, X4, X5]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))&(m1_pre_topc(X3,X1)&(m1_pre_topc(X4,X1)&(v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))))))=>(v1_funct_1(k3_tmap_1(X1,X2,X3,X4,X5))&(v1_funct_2(k3_tmap_1(X1,X2,X3,X4,X5),u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(k3_tmap_1(X1,X2,X3,X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))), file('data/MPTP/proofs/66df56672ed048759d4dbc07d81addf0.E_input', dt_k3_tmap_1)).
fof(t1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/66df56672ed048759d4dbc07d81addf0.E_input', t1_tsep_1)).
fof(t2_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>m1_pre_topc(X1,X1)), file('data/MPTP/proofs/66df56672ed048759d4dbc07d81addf0.E_input', t2_tsep_1)).
fof(t55_pre_topc, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>m1_subset_1(X3,u1_struct_0(X1))))), file('data/MPTP/proofs/66df56672ed048759d4dbc07d81addf0.E_input', t55_pre_topc)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/66df56672ed048759d4dbc07d81addf0.E_input', t5_subset)).
fof(d5_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:(m1_pre_topc(X3,X1)=>![X4]:(m1_pre_topc(X4,X1)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))=>(m1_pre_topc(X4,X3)=>k3_tmap_1(X1,X2,X3,X4,X5)=k2_partfun1(u1_struct_0(X3),u1_struct_0(X2),X5,u1_struct_0(X4)))))))), file('data/MPTP/proofs/66df56672ed048759d4dbc07d81addf0.E_input', d5_tmap_1)).
fof(t7_tsep_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X2)=>m1_pre_topc(X3,X1)))), file('data/MPTP/proofs/66df56672ed048759d4dbc07d81addf0.E_input', t7_tsep_1)).
fof(d4_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:(m1_pre_topc(X4,X1)=>k2_tmap_1(X1,X2,X3,X4)=k2_partfun1(u1_struct_0(X1),u1_struct_0(X2),X3,u1_struct_0(X4)))))), file('data/MPTP/proofs/66df56672ed048759d4dbc07d81addf0.E_input', d4_tmap_1)).
fof(t74_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))=>r2_funct_2(u1_struct_0(X3),u1_struct_0(X2),X4,k3_tmap_1(X1,X2,X3,X3,X4)))))), file('data/MPTP/proofs/66df56672ed048759d4dbc07d81addf0.E_input', t74_tmap_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/66df56672ed048759d4dbc07d81addf0.E_input', t2_subset)).
fof(redefinition_r2_funct_2, axiom, ![X1, X2, X3, X4]:(((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&(v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))=>(r2_funct_2(X1,X2,X3,X4)<=>X3=X4)), file('data/MPTP/proofs/66df56672ed048759d4dbc07d81addf0.E_input', redefinition_r2_funct_2)).
fof(t95_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,u1_struct_0(X2))=>k1_funct_1(k4_tmap_1(X1,X2),X3)=X3)))), file('data/MPTP/proofs/66df56672ed048759d4dbc07d81addf0.E_input', t95_tmap_1)).
fof(c_0_18, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,u1_struct_0(X2))=>X4=k1_funct_1(X3,X4)))=>r2_funct_2(u1_struct_0(X2),u1_struct_0(X1),k4_tmap_1(X1,X2),X3)))))), inference(assume_negation,[status(cth)],[t96_tmap_1])).
fof(c_0_19, plain, ![X11, X12]:((v1_funct_1(k4_tmap_1(X11,X12))|(v2_struct_0(X11)|(~v2_pre_topc(X11)|~l1_pre_topc(X11))|~m1_pre_topc(X12,X11)))&((v1_funct_2(k4_tmap_1(X11,X12),u1_struct_0(X12),u1_struct_0(X11))|(v2_struct_0(X11)|(~v2_pre_topc(X11)|~l1_pre_topc(X11))|~m1_pre_topc(X12,X11)))&(m1_subset_1(k4_tmap_1(X11,X12),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X12),u1_struct_0(X11))))|(v2_struct_0(X11)|(~v2_pre_topc(X11)|~l1_pre_topc(X11))|~m1_pre_topc(X12,X11))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_tmap_1])])])])).
fof(c_0_20, negated_conjecture, ![X10]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,u1_struct_0(esk2_0),u1_struct_0(esk1_0))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))))&((~m1_subset_1(X10,u1_struct_0(esk1_0))|(~r2_hidden(X10,u1_struct_0(esk2_0))|X10=k1_funct_1(esk3_0,X10)))&~r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_18])])])])])).
fof(c_0_21, plain, ![X13, X14, X15, X16, X17]:((~r2_funct_2(X13,X14,X15,X16)|(~m1_subset_1(X17,X13)|k1_funct_1(X15,X17)=k1_funct_1(X16,X17))|(~v1_funct_1(X16)|(~v1_funct_2(X16,X13,X14)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X13,X14)))))|(~v1_funct_1(X15)|(~v1_funct_2(X15,X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))))&((m1_subset_1(esk4_4(X13,X14,X15,X16),X13)|r2_funct_2(X13,X14,X15,X16)|(~v1_funct_1(X16)|(~v1_funct_2(X16,X13,X14)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X13,X14)))))|(~v1_funct_1(X15)|(~v1_funct_2(X15,X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))))&(k1_funct_1(X15,esk4_4(X13,X14,X15,X16))!=k1_funct_1(X16,esk4_4(X13,X14,X15,X16))|r2_funct_2(X13,X14,X15,X16)|(~v1_funct_1(X16)|(~v1_funct_2(X16,X13,X14)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X13,X14)))))|(~v1_funct_1(X15)|(~v1_funct_2(X15,X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_funct_2])])])])])).
fof(c_0_22, plain, ![X19, X20, X21, X22, X23]:((m1_subset_1(esk5_5(X19,X20,X21,X22,X23),u1_struct_0(X20))|r2_funct_2(u1_struct_0(X21),u1_struct_0(X19),k2_tmap_1(X20,X19,X22,X21),X23)|(~v1_funct_1(X23)|(~v1_funct_2(X23,u1_struct_0(X21),u1_struct_0(X19))|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X21),u1_struct_0(X19))))))|(~v1_funct_1(X22)|(~v1_funct_2(X22,u1_struct_0(X20),u1_struct_0(X19))|~m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X20),u1_struct_0(X19))))))|(v2_struct_0(X21)|~m1_pre_topc(X21,X20))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20)))|(v2_struct_0(X19)|(~v2_pre_topc(X19)|~l1_pre_topc(X19))))&((r2_hidden(esk5_5(X19,X20,X21,X22,X23),u1_struct_0(X21))|r2_funct_2(u1_struct_0(X21),u1_struct_0(X19),k2_tmap_1(X20,X19,X22,X21),X23)|(~v1_funct_1(X23)|(~v1_funct_2(X23,u1_struct_0(X21),u1_struct_0(X19))|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X21),u1_struct_0(X19))))))|(~v1_funct_1(X22)|(~v1_funct_2(X22,u1_struct_0(X20),u1_struct_0(X19))|~m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X20),u1_struct_0(X19))))))|(v2_struct_0(X21)|~m1_pre_topc(X21,X20))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20)))|(v2_struct_0(X19)|(~v2_pre_topc(X19)|~l1_pre_topc(X19))))&(k3_funct_2(u1_struct_0(X20),u1_struct_0(X19),X22,esk5_5(X19,X20,X21,X22,X23))!=k1_funct_1(X23,esk5_5(X19,X20,X21,X22,X23))|r2_funct_2(u1_struct_0(X21),u1_struct_0(X19),k2_tmap_1(X20,X19,X22,X21),X23)|(~v1_funct_1(X23)|(~v1_funct_2(X23,u1_struct_0(X21),u1_struct_0(X19))|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X21),u1_struct_0(X19))))))|(~v1_funct_1(X22)|(~v1_funct_2(X22,u1_struct_0(X20),u1_struct_0(X19))|~m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X20),u1_struct_0(X19))))))|(v2_struct_0(X21)|~m1_pre_topc(X21,X20))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20)))|(v2_struct_0(X19)|(~v2_pre_topc(X19)|~l1_pre_topc(X19)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t59_tmap_1])])])])])])).
fof(c_0_23, plain, ![X25, X26]:(~l1_pre_topc(X25)|(~m1_pre_topc(X26,X25)|l1_pre_topc(X26))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
fof(c_0_24, plain, ![X27, X28]:(~v2_pre_topc(X27)|~l1_pre_topc(X27)|(~m1_pre_topc(X28,X27)|v2_pre_topc(X28))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_pre_topc])])])).
fof(c_0_25, plain, ![X29, X30, X31, X32, X33]:((v1_funct_1(k3_tmap_1(X29,X30,X31,X32,X33))|(v2_struct_0(X29)|(~v2_pre_topc(X29)|~l1_pre_topc(X29))|(v2_struct_0(X30)|(~v2_pre_topc(X30)|~l1_pre_topc(X30))|(~m1_pre_topc(X31,X29)|(~m1_pre_topc(X32,X29)|(~v1_funct_1(X33)|(~v1_funct_2(X33,u1_struct_0(X31),u1_struct_0(X30))|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X31),u1_struct_0(X30)))))))))))&((v1_funct_2(k3_tmap_1(X29,X30,X31,X32,X33),u1_struct_0(X32),u1_struct_0(X30))|(v2_struct_0(X29)|(~v2_pre_topc(X29)|~l1_pre_topc(X29))|(v2_struct_0(X30)|(~v2_pre_topc(X30)|~l1_pre_topc(X30))|(~m1_pre_topc(X31,X29)|(~m1_pre_topc(X32,X29)|(~v1_funct_1(X33)|(~v1_funct_2(X33,u1_struct_0(X31),u1_struct_0(X30))|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X31),u1_struct_0(X30)))))))))))&(m1_subset_1(k3_tmap_1(X29,X30,X31,X32,X33),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X32),u1_struct_0(X30))))|(v2_struct_0(X29)|(~v2_pre_topc(X29)|~l1_pre_topc(X29))|(v2_struct_0(X30)|(~v2_pre_topc(X30)|~l1_pre_topc(X30))|(~m1_pre_topc(X31,X29)|(~m1_pre_topc(X32,X29)|(~v1_funct_1(X33)|(~v1_funct_2(X33,u1_struct_0(X31),u1_struct_0(X30))|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X31),u1_struct_0(X30))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_tmap_1])])])])).
cnf(c_0_26, plain, (m1_subset_1(k4_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_31, plain, (v1_funct_2(k4_tmap_1(X1,X2),u1_struct_0(X2),u1_struct_0(X1))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_32, plain, (v1_funct_1(k4_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_33, plain, (m1_subset_1(esk4_4(X1,X2,X3,X4),X1)|r2_funct_2(X1,X2,X3,X4)|~v1_funct_1(X4)|~v1_funct_2(X4,X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_34, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_35, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk2_0),u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_36, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_37, plain, ![X34, X35]:(~l1_pre_topc(X34)|(~m1_pre_topc(X35,X34)|m1_subset_1(u1_struct_0(X35),k1_zfmisc_1(u1_struct_0(X34))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])).
fof(c_0_38, plain, ![X36]:(~l1_pre_topc(X36)|m1_pre_topc(X36,X36)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tsep_1])])).
cnf(c_0_39, plain, (r2_hidden(esk5_5(X1,X2,X3,X4,X5),u1_struct_0(X3))|r2_funct_2(u1_struct_0(X3),u1_struct_0(X1),k2_tmap_1(X2,X1,X4,X3),X5)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X1))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_1(X4)|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_pre_topc(X3,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_40, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_41, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_42, plain, (v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_43, plain, (m1_subset_1(k3_tmap_1(X1,X2,X3,X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_44, negated_conjecture, (m1_subset_1(k4_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_29])]), c_0_30])).
cnf(c_0_45, negated_conjecture, (v1_funct_2(k4_tmap_1(esk1_0,esk2_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_27]), c_0_28]), c_0_29])]), c_0_30])).
cnf(c_0_46, negated_conjecture, (v1_funct_1(k4_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_27]), c_0_28]), c_0_29])]), c_0_30])).
cnf(c_0_47, plain, (v1_funct_2(k3_tmap_1(X1,X2,X3,X4,X5),u1_struct_0(X4),u1_struct_0(X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_48, plain, (v1_funct_1(k3_tmap_1(X1,X2,X3,X4,X5))|v2_struct_0(X1)|v2_struct_0(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_49, plain, ![X37, X38, X39]:(v2_struct_0(X37)|~l1_pre_topc(X37)|(v2_struct_0(X38)|~m1_pre_topc(X38,X37)|(~m1_subset_1(X39,u1_struct_0(X38))|m1_subset_1(X39,u1_struct_0(X37))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t55_pre_topc])])])])).
cnf(c_0_50, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),X1,esk3_0)|m1_subset_1(esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),X1,esk3_0),u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35]), c_0_36])])).
cnf(c_0_51, negated_conjecture, (~r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_52, plain, ![X40, X41, X42]:(~r2_hidden(X40,X41)|(~m1_subset_1(X41,k1_zfmisc_1(X42))|~v1_xboole_0(X42))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
cnf(c_0_53, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_54, plain, (m1_pre_topc(X1,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_55, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tmap_1(X1,esk1_0,X2,esk2_0),esk3_0)|r2_hidden(esk5_5(esk1_0,X1,esk2_0,X2,esk3_0),u1_struct_0(esk2_0))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_34]), c_0_35]), c_0_36]), c_0_28]), c_0_29])]), c_0_30]), c_0_40])).
cnf(c_0_56, negated_conjecture, (l1_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_27]), c_0_28])])).
cnf(c_0_57, negated_conjecture, (v2_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_27]), c_0_28]), c_0_29])])).
fof(c_0_58, plain, ![X43, X44, X45, X46, X47]:(v2_struct_0(X43)|(~v2_pre_topc(X43)|~l1_pre_topc(X43))|(v2_struct_0(X44)|(~v2_pre_topc(X44)|~l1_pre_topc(X44))|(~m1_pre_topc(X45,X43)|(~m1_pre_topc(X46,X43)|(~v1_funct_1(X47)|(~v1_funct_2(X47,u1_struct_0(X45),u1_struct_0(X44))|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X45),u1_struct_0(X44)))))|(~m1_pre_topc(X46,X45)|k3_tmap_1(X43,X44,X45,X46,X47)=k2_partfun1(u1_struct_0(X45),u1_struct_0(X44),X47,u1_struct_0(X46)))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_tmap_1])])])])).
fof(c_0_59, plain, ![X48, X49, X50]:(~v2_pre_topc(X48)|~l1_pre_topc(X48)|(~m1_pre_topc(X49,X48)|(~m1_pre_topc(X50,X49)|m1_pre_topc(X50,X48)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tsep_1])])])).
fof(c_0_60, plain, ![X51, X52, X53, X54]:(v2_struct_0(X51)|(~v2_pre_topc(X51)|~l1_pre_topc(X51))|(v2_struct_0(X52)|(~v2_pre_topc(X52)|~l1_pre_topc(X52))|(~v1_funct_1(X53)|(~v1_funct_2(X53,u1_struct_0(X51),u1_struct_0(X52))|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X51),u1_struct_0(X52)))))|(~m1_pre_topc(X54,X51)|k2_tmap_1(X51,X52,X53,X54)=k2_partfun1(u1_struct_0(X51),u1_struct_0(X52),X53,u1_struct_0(X54)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_tmap_1])])])])).
cnf(c_0_61, negated_conjecture, (m1_subset_1(k3_tmap_1(X1,esk1_0,esk2_0,X2,k4_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|v2_struct_0(X1)|~m1_pre_topc(esk2_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_46]), c_0_28]), c_0_29])]), c_0_30])).
cnf(c_0_62, negated_conjecture, (v1_funct_2(k3_tmap_1(X1,esk1_0,esk2_0,X2,k4_tmap_1(esk1_0,esk2_0)),u1_struct_0(X2),u1_struct_0(esk1_0))|v2_struct_0(X1)|~m1_pre_topc(esk2_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_44]), c_0_45]), c_0_46]), c_0_28]), c_0_29])]), c_0_30])).
cnf(c_0_63, negated_conjecture, (v1_funct_1(k3_tmap_1(X1,esk1_0,esk2_0,X2,k4_tmap_1(esk1_0,esk2_0)))|v2_struct_0(X1)|~m1_pre_topc(esk2_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_44]), c_0_45]), c_0_46]), c_0_28]), c_0_29])]), c_0_30])).
fof(c_0_64, plain, ![X55, X56, X57, X58]:(v2_struct_0(X55)|(~v2_pre_topc(X55)|~l1_pre_topc(X55))|(v2_struct_0(X56)|(~v2_pre_topc(X56)|~l1_pre_topc(X56))|(v2_struct_0(X57)|~m1_pre_topc(X57,X55)|(~v1_funct_1(X58)|(~v1_funct_2(X58,u1_struct_0(X57),u1_struct_0(X56))|~m1_subset_1(X58,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X57),u1_struct_0(X56)))))|r2_funct_2(u1_struct_0(X57),u1_struct_0(X56),X58,k3_tmap_1(X55,X56,X57,X57,X58)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t74_tmap_1])])])])).
cnf(c_0_65, plain, (v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(X3,u1_struct_0(X1))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_66, negated_conjecture, (m1_subset_1(esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_44]), c_0_45]), c_0_46])]), c_0_51])).
cnf(c_0_67, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_68, plain, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_69, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tmap_1(esk2_0,esk1_0,k4_tmap_1(esk1_0,esk2_0),esk2_0),esk3_0)|r2_hidden(esk5_5(esk1_0,esk2_0,esk2_0,k4_tmap_1(esk1_0,esk2_0),esk3_0),u1_struct_0(esk2_0))|~m1_pre_topc(esk2_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_44]), c_0_45]), c_0_46]), c_0_56]), c_0_57])]), c_0_40])).
fof(c_0_70, plain, ![X59, X60]:(~m1_subset_1(X59,X60)|(v1_xboole_0(X60)|r2_hidden(X59,X60))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_71, plain, (v2_struct_0(X1)|v2_struct_0(X2)|k3_tmap_1(X1,X2,X3,X4,X5)=k2_partfun1(u1_struct_0(X3),u1_struct_0(X2),X5,u1_struct_0(X4))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_pre_topc(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_72, plain, (m1_pre_topc(X3,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_73, plain, (v2_struct_0(X1)|v2_struct_0(X2)|k2_tmap_1(X1,X2,X3,X4)=k2_partfun1(u1_struct_0(X1),u1_struct_0(X2),X3,u1_struct_0(X4))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_pre_topc(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
fof(c_0_74, plain, ![X61, X62, X63, X64]:((~r2_funct_2(X61,X62,X63,X64)|X63=X64|(~v1_funct_1(X63)|(~v1_funct_2(X63,X61,X62)|~m1_subset_1(X63,k1_zfmisc_1(k2_zfmisc_1(X61,X62))))|(~v1_funct_1(X64)|(~v1_funct_2(X64,X61,X62)|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(X61,X62)))))))&(X63!=X64|r2_funct_2(X61,X62,X63,X64)|(~v1_funct_1(X63)|(~v1_funct_2(X63,X61,X62)|~m1_subset_1(X63,k1_zfmisc_1(k2_zfmisc_1(X61,X62))))|(~v1_funct_1(X64)|(~v1_funct_2(X64,X61,X62)|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(X61,X62)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r2_funct_2])])])).
cnf(c_0_75, negated_conjecture, (m1_subset_1(k3_tmap_1(X1,esk1_0,esk2_0,X1,k4_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|v2_struct_0(X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_61, c_0_54])).
cnf(c_0_76, negated_conjecture, (v1_funct_2(k3_tmap_1(X1,esk1_0,esk2_0,X1,k4_tmap_1(esk1_0,esk2_0)),u1_struct_0(X1),u1_struct_0(esk1_0))|v2_struct_0(X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_62, c_0_54])).
cnf(c_0_77, negated_conjecture, (v1_funct_1(k3_tmap_1(X1,esk1_0,esk2_0,X1,k4_tmap_1(esk1_0,esk2_0)))|v2_struct_0(X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_63, c_0_54])).
cnf(c_0_78, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|r2_funct_2(u1_struct_0(X3),u1_struct_0(X2),X4,k3_tmap_1(X1,X2,X3,X3,X4))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~v1_funct_1(X4)|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_79, negated_conjecture, (m1_subset_1(esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0),u1_struct_0(X1))|v2_struct_0(X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_40])).
cnf(c_0_80, plain, (~v1_xboole_0(u1_struct_0(X1))|~r2_hidden(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_81, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tmap_1(esk2_0,esk1_0,k4_tmap_1(esk1_0,esk2_0),esk2_0),esk3_0)|r2_hidden(esk5_5(esk1_0,esk2_0,esk2_0,k4_tmap_1(esk1_0,esk2_0),esk3_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_54]), c_0_56])])).
cnf(c_0_82, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_83, plain, (k3_tmap_1(X1,X2,X3,X4,X5)=k2_partfun1(u1_struct_0(X3),u1_struct_0(X2),X5,u1_struct_0(X4))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_pre_topc(X4,X3)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_84, negated_conjecture, (k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),u1_struct_0(X1))=k2_tmap_1(esk2_0,esk1_0,k4_tmap_1(esk1_0,esk2_0),X1)|~m1_pre_topc(X1,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_44]), c_0_45]), c_0_46]), c_0_28]), c_0_56]), c_0_29]), c_0_57])]), c_0_40]), c_0_30])).
cnf(c_0_85, plain, (X3=X4|~r2_funct_2(X1,X2,X3,X4)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X4)|~v1_funct_2(X4,X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_86, negated_conjecture, (m1_subset_1(k3_tmap_1(esk2_0,esk1_0,esk2_0,esk2_0,k4_tmap_1(esk1_0,esk2_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_54]), c_0_56]), c_0_57])]), c_0_40])).
cnf(c_0_87, negated_conjecture, (v1_funct_2(k3_tmap_1(esk2_0,esk1_0,esk2_0,esk2_0,k4_tmap_1(esk1_0,esk2_0)),u1_struct_0(esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_54]), c_0_56]), c_0_57])]), c_0_40])).
cnf(c_0_88, negated_conjecture, (v1_funct_1(k3_tmap_1(esk2_0,esk1_0,esk2_0,esk2_0,k4_tmap_1(esk1_0,esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_54]), c_0_56]), c_0_57])]), c_0_40])).
cnf(c_0_89, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),k3_tmap_1(X1,esk1_0,esk2_0,esk2_0,k4_tmap_1(esk1_0,esk2_0)))|v2_struct_0(X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_44]), c_0_45]), c_0_46]), c_0_28]), c_0_29])]), c_0_30]), c_0_40])).
cnf(c_0_90, negated_conjecture, (X1=k1_funct_1(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_91, negated_conjecture, (m1_subset_1(esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_27]), c_0_28])]), c_0_30])).
cnf(c_0_92, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tmap_1(esk2_0,esk1_0,k4_tmap_1(esk1_0,esk2_0),esk2_0),esk3_0)|~v1_xboole_0(u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_56])])).
cnf(c_0_93, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|r2_hidden(esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_82, c_0_66])).
cnf(c_0_94, negated_conjecture, (k3_tmap_1(X1,esk1_0,esk2_0,X2,k4_tmap_1(esk1_0,esk2_0))=k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),u1_struct_0(X2))|v2_struct_0(X1)|~m1_pre_topc(X2,esk2_0)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_44]), c_0_45]), c_0_46]), c_0_28]), c_0_29])]), c_0_30])).
cnf(c_0_95, negated_conjecture, (k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),u1_struct_0(esk2_0))=k2_tmap_1(esk2_0,esk1_0,k4_tmap_1(esk1_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_54]), c_0_56])])).
cnf(c_0_96, negated_conjecture, (X1=k3_tmap_1(esk2_0,esk1_0,esk2_0,esk2_0,k4_tmap_1(esk1_0,esk2_0))|~r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),X1,k3_tmap_1(esk2_0,esk1_0,esk2_0,esk2_0,k4_tmap_1(esk1_0,esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_87]), c_0_88])])).
cnf(c_0_97, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),k3_tmap_1(esk2_0,esk1_0,esk2_0,esk2_0,k4_tmap_1(esk1_0,esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_54]), c_0_56]), c_0_57])]), c_0_40])).
cnf(c_0_98, plain, (r2_funct_2(X2,X3,X1,X4)|k1_funct_1(X1,esk4_4(X2,X3,X1,X4))!=k1_funct_1(X4,esk4_4(X2,X3,X1,X4))|~v1_funct_1(X4)|~v1_funct_2(X4,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_99, negated_conjecture, (k1_funct_1(esk3_0,esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0))=esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0)|~r2_hidden(esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_100, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tmap_1(esk2_0,esk1_0,k4_tmap_1(esk1_0,esk2_0),esk2_0),esk3_0)|r2_hidden(esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_101, negated_conjecture, (k3_tmap_1(X1,esk1_0,esk2_0,esk2_0,k4_tmap_1(esk1_0,esk2_0))=k2_tmap_1(esk2_0,esk1_0,k4_tmap_1(esk1_0,esk2_0),esk2_0)|v2_struct_0(X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_54]), c_0_95]), c_0_56])])).
cnf(c_0_102, negated_conjecture, (k3_tmap_1(esk2_0,esk1_0,esk2_0,esk2_0,k4_tmap_1(esk1_0,esk2_0))=k4_tmap_1(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_44]), c_0_97]), c_0_45]), c_0_46])])).
fof(c_0_103, plain, ![X65, X66, X67]:(v2_struct_0(X65)|(~v2_pre_topc(X65)|~l1_pre_topc(X65))|(v2_struct_0(X66)|~m1_pre_topc(X66,X65)|(~m1_subset_1(X67,u1_struct_0(X65))|(~r2_hidden(X67,u1_struct_0(X66))|k1_funct_1(k4_tmap_1(X65,X66),X67)=X67)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t95_tmap_1])])])])).
cnf(c_0_104, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),X1,esk3_0)|k1_funct_1(X1,esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),X1,esk3_0))!=k1_funct_1(esk3_0,esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),X1,esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_34]), c_0_35]), c_0_36])])).
cnf(c_0_105, negated_conjecture, (k1_funct_1(esk3_0,esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0))=esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0)|r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tmap_1(esk2_0,esk1_0,k4_tmap_1(esk1_0,esk2_0),esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_106, negated_conjecture, (k2_tmap_1(esk2_0,esk1_0,k4_tmap_1(esk1_0,esk2_0),esk2_0)=k4_tmap_1(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_54]), c_0_102]), c_0_56]), c_0_57])]), c_0_40])).
cnf(c_0_107, plain, (v2_struct_0(X1)|v2_struct_0(X2)|k1_funct_1(k4_tmap_1(X1,X2),X3)=X3|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_108, negated_conjecture, (k1_funct_1(k4_tmap_1(esk1_0,esk2_0),esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0))!=k1_funct_1(esk3_0,esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_44]), c_0_45]), c_0_46])]), c_0_51])).
cnf(c_0_109, negated_conjecture, (k1_funct_1(esk3_0,esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0))=esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_105, c_0_106]), c_0_51])).
cnf(c_0_110, negated_conjecture, (k1_funct_1(k4_tmap_1(esk1_0,X1),esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0))=esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0)|v2_struct_0(X1)|~r2_hidden(esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0),u1_struct_0(X1))|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_91]), c_0_28]), c_0_29])]), c_0_30])).
cnf(c_0_111, negated_conjecture, (r2_hidden(esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_100, c_0_106]), c_0_51])).
cnf(c_0_112, negated_conjecture, (k1_funct_1(k4_tmap_1(esk1_0,esk2_0),esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0))!=esk4_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0)), inference(rw,[status(thm)],[c_0_108, c_0_109])).
cnf(c_0_113, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_27]), c_0_111])]), c_0_112]), c_0_40]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 114
# Proof object clause steps            : 77
# Proof object formula steps           : 37
# Proof object conjectures             : 55
# Proof object clause conjectures      : 52
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 32
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 41
# Proof object simplifying inferences  : 139
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 36
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 36
# Processed clauses                    : 662
# ...of these trivial                  : 6
# ...subsumed                          : 80
# ...remaining for further processing  : 576
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 31
# Backward-rewritten                   : 132
# Generated clauses                    : 1113
# ...of the previous two non-trivial   : 1123
# Contextual simplify-reflections      : 44
# Paramodulations                      : 1108
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 408
#    Positive orientable unit clauses  : 125
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 276
# Current number of unprocessed clauses: 435
# ...number of literals in the above   : 2161
# Current number of archived formulas  : 0
# Current number of archived clauses   : 167
# Clause-clause subsumption calls (NU) : 13364
# Rec. Clause-clause subsumption calls : 1390
# Non-unit clause-clause subsumptions  : 126
# Unit Clause-clause subsumption calls : 907
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 758
# BW rewrite match successes           : 38
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 76408

# -------------------------------------------------
# User time                : 0.049 s
# System time              : 0.000 s
# Total time               : 0.049 s
# Maximum resident set size: 28748 pages
