# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_AE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t50_funct_1, conjecture, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X2)=X1&(k1_relat_1(X3)=X1&(r1_tarski(k2_relat_1(X3),X1)&(v2_funct_1(X2)&k5_relat_1(X3,X2)=X2))))=>X3=k6_relat_1(X1)))), file('data/MPTP/proofs/67061756a90a42fdb40209885704dc23.E_input', t50_funct_1)).
fof(t49_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)<=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((r1_tarski(k2_relat_1(X2),k1_relat_1(X1))&(r1_tarski(k2_relat_1(X3),k1_relat_1(X1))&(k1_relat_1(X2)=k1_relat_1(X3)&k5_relat_1(X2,X1)=k5_relat_1(X3,X1))))=>X2=X3))))), file('data/MPTP/proofs/67061756a90a42fdb40209885704dc23.E_input', t49_funct_1)).
fof(t71_relat_1, axiom, ![X1]:(k1_relat_1(k6_relat_1(X1))=X1&k2_relat_1(k6_relat_1(X1))=X1), file('data/MPTP/proofs/67061756a90a42fdb40209885704dc23.E_input', t71_relat_1)).
fof(fc3_funct_1, axiom, ![X1]:(v1_relat_1(k6_relat_1(X1))&v1_funct_1(k6_relat_1(X1))), file('data/MPTP/proofs/67061756a90a42fdb40209885704dc23.E_input', fc3_funct_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/67061756a90a42fdb40209885704dc23.E_input', d3_tarski)).
fof(t78_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k5_relat_1(k6_relat_1(k1_relat_1(X1)),X1)=X1), file('data/MPTP/proofs/67061756a90a42fdb40209885704dc23.E_input', t78_relat_1)).
fof(c_0_6, negated_conjecture, ~(![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X2)=X1&(k1_relat_1(X3)=X1&(r1_tarski(k2_relat_1(X3),X1)&(v2_funct_1(X2)&k5_relat_1(X3,X2)=X2))))=>X3=k6_relat_1(X1))))), inference(assume_negation,[status(cth)],[t50_funct_1])).
fof(c_0_7, plain, ![X7, X8, X9]:((~v2_funct_1(X7)|(~v1_relat_1(X8)|~v1_funct_1(X8)|(~v1_relat_1(X9)|~v1_funct_1(X9)|(~r1_tarski(k2_relat_1(X8),k1_relat_1(X7))|(~r1_tarski(k2_relat_1(X9),k1_relat_1(X7))|(k1_relat_1(X8)!=k1_relat_1(X9)|k5_relat_1(X8,X7)!=k5_relat_1(X9,X7)))|X8=X9)))|(~v1_relat_1(X7)|~v1_funct_1(X7)))&(((v1_relat_1(esk4_1(X7))|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&(v1_funct_1(esk4_1(X7))|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7))))&(((v1_relat_1(esk5_1(X7))|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&(v1_funct_1(esk5_1(X7))|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7))))&(((r1_tarski(k2_relat_1(esk4_1(X7)),k1_relat_1(X7))|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&((r1_tarski(k2_relat_1(esk5_1(X7)),k1_relat_1(X7))|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&((k1_relat_1(esk4_1(X7))=k1_relat_1(esk5_1(X7))|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&(k5_relat_1(esk4_1(X7),X7)=k5_relat_1(esk5_1(X7),X7)|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7))))))&(esk4_1(X7)!=esk5_1(X7)|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t49_funct_1])])])])])).
fof(c_0_8, negated_conjecture, ((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&((v1_relat_1(esk3_0)&v1_funct_1(esk3_0))&((k1_relat_1(esk2_0)=esk1_0&(k1_relat_1(esk3_0)=esk1_0&(r1_tarski(k2_relat_1(esk3_0),esk1_0)&(v2_funct_1(esk2_0)&k5_relat_1(esk3_0,esk2_0)=esk2_0))))&esk3_0!=k6_relat_1(esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
cnf(c_0_9, plain, (X2=X3|~v2_funct_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_funct_1(X3)|~r1_tarski(k2_relat_1(X2),k1_relat_1(X1))|~r1_tarski(k2_relat_1(X3),k1_relat_1(X1))|k1_relat_1(X2)!=k1_relat_1(X3)|k5_relat_1(X2,X1)!=k5_relat_1(X3,X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10, negated_conjecture, (k1_relat_1(esk2_0)=esk1_0), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11, negated_conjecture, (v2_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, negated_conjecture, (X1=X2|k5_relat_1(X1,esk2_0)!=k5_relat_1(X2,esk2_0)|k1_relat_1(X1)!=k1_relat_1(X2)|~r1_tarski(k2_relat_1(X2),esk1_0)|~r1_tarski(k2_relat_1(X1),esk1_0)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11]), c_0_12]), c_0_13])])).
cnf(c_0_15, negated_conjecture, (r1_tarski(k2_relat_1(esk3_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16, negated_conjecture, (k5_relat_1(esk3_0,esk2_0)=esk2_0), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_17, negated_conjecture, (k1_relat_1(esk3_0)=esk1_0), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_18, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_19, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_20, plain, ![X12]:(k1_relat_1(k6_relat_1(X12))=X12&k2_relat_1(k6_relat_1(X12))=X12), inference(variable_rename,[status(thm)],[t71_relat_1])).
fof(c_0_21, plain, ![X13]:(v1_relat_1(k6_relat_1(X13))&v1_funct_1(k6_relat_1(X13))), inference(variable_rename,[status(thm)],[fc3_funct_1])).
fof(c_0_22, plain, ![X14, X15, X16, X17, X18]:((~r1_tarski(X14,X15)|(~r2_hidden(X16,X14)|r2_hidden(X16,X15)))&((r2_hidden(esk6_2(X17,X18),X17)|r1_tarski(X17,X18))&(~r2_hidden(esk6_2(X17,X18),X18)|r1_tarski(X17,X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_23, negated_conjecture, (X1=esk3_0|k5_relat_1(X1,esk2_0)!=esk2_0|k1_relat_1(X1)!=esk1_0|~r1_tarski(k2_relat_1(X1),esk1_0)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19])])).
cnf(c_0_24, plain, (k2_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25, plain, (k1_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (v1_funct_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_28, plain, ![X20]:(~v1_relat_1(X20)|k5_relat_1(k6_relat_1(k1_relat_1(X20)),X20)=X20), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t78_relat_1])])).
cnf(c_0_29, plain, (r1_tarski(X1,X2)|~r2_hidden(esk6_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, negated_conjecture, (k6_relat_1(X1)=esk3_0|k5_relat_1(k6_relat_1(X1),esk2_0)!=esk2_0|X1!=esk1_0|~r1_tarski(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25]), c_0_26]), c_0_27])])).
cnf(c_0_32, plain, (k5_relat_1(k6_relat_1(k1_relat_1(X1)),X1)=X1|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_34, negated_conjecture, (esk3_0!=k6_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_35, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_10]), c_0_10]), c_0_10]), c_0_33]), c_0_13])]), c_0_34]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 36
# Proof object clause steps            : 23
# Proof object formula steps           : 13
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 20
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 28
# Processed clauses                    : 52
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 50
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 53
# ...of the previous two non-trivial   : 45
# Contextual simplify-reflections      : 1
# Paramodulations                      : 53
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 50
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 32
# Current number of unprocessed clauses: 21
# ...number of literals in the above   : 166
# Current number of archived formulas  : 0
# Current number of archived clauses   : 0
# Clause-clause subsumption calls (NU) : 173
# Rec. Clause-clause subsumption calls : 82
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2870

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.004 s
# Total time               : 0.010 s
# Maximum resident set size: 28712 pages
