# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t22_relset_1, conjecture, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>(![X4]:~((r2_hidden(X4,X2)&![X5]:~(r2_hidden(k4_tarski(X4,X5),X3))))<=>k1_relset_1(X2,X1,X3)=X2)), file('data/MPTP/proofs/672ae90fcf7e433ca78a0be8e7778fb0.E_input', t22_relset_1)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/672ae90fcf7e433ca78a0be8e7778fb0.E_input', l3_subset_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/672ae90fcf7e433ca78a0be8e7778fb0.E_input', redefinition_k1_relset_1)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/672ae90fcf7e433ca78a0be8e7778fb0.E_input', l54_zfmisc_1)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/672ae90fcf7e433ca78a0be8e7778fb0.E_input', d4_relat_1)).
fof(c_0_5, negated_conjecture, ~(![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>(![X4]:~((r2_hidden(X4,X2)&![X5]:~(r2_hidden(k4_tarski(X4,X5),X3))))<=>k1_relset_1(X2,X1,X3)=X2))), inference(assume_negation,[status(cth)],[t22_relset_1])).
fof(c_0_6, plain, ![X13, X14, X15]:(~m1_subset_1(X14,k1_zfmisc_1(X13))|(~r2_hidden(X15,X14)|r2_hidden(X15,X13))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
fof(c_0_7, negated_conjecture, ![X10, X11]:(m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))&(((r2_hidden(esk4_0,esk2_0)|k1_relset_1(esk2_0,esk1_0,esk3_0)!=esk2_0)&(~r2_hidden(k4_tarski(esk4_0,X10),esk3_0)|k1_relset_1(esk2_0,esk1_0,esk3_0)!=esk2_0))&(~r2_hidden(X11,esk2_0)|r2_hidden(k4_tarski(X11,esk5_1(X11)),esk3_0)|k1_relset_1(esk2_0,esk1_0,esk3_0)=esk2_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_5])])])])])])).
fof(c_0_8, plain, ![X16, X17, X18]:(~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(X16,X17)))|k1_relset_1(X16,X17,X18)=k1_relat_1(X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
fof(c_0_9, plain, ![X19, X20, X21, X22]:(((r2_hidden(X19,X21)|~r2_hidden(k4_tarski(X19,X20),k2_zfmisc_1(X21,X22)))&(r2_hidden(X20,X22)|~r2_hidden(k4_tarski(X19,X20),k2_zfmisc_1(X21,X22))))&(~r2_hidden(X19,X21)|~r2_hidden(X20,X22)|r2_hidden(k4_tarski(X19,X20),k2_zfmisc_1(X21,X22)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
cnf(c_0_10, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(esk2_0,esk1_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
fof(c_0_15, plain, ![X23, X24, X25, X27, X28, X29, X30, X32]:(((~r2_hidden(X25,X24)|r2_hidden(k4_tarski(X25,esk6_3(X23,X24,X25)),X23)|X24!=k1_relat_1(X23))&(~r2_hidden(k4_tarski(X27,X28),X23)|r2_hidden(X27,X24)|X24!=k1_relat_1(X23)))&((~r2_hidden(esk7_2(X29,X30),X30)|~r2_hidden(k4_tarski(esk7_2(X29,X30),X32),X29)|X30=k1_relat_1(X29))&(r2_hidden(esk7_2(X29,X30),X30)|r2_hidden(k4_tarski(esk7_2(X29,X30),esk8_2(X29,X30)),X29)|X30=k1_relat_1(X29)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
cnf(c_0_16, negated_conjecture, (r2_hidden(k4_tarski(X1,esk5_1(X1)),esk3_0)|k1_relset_1(esk2_0,esk1_0,esk3_0)=esk2_0|~r2_hidden(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_17, negated_conjecture, (k1_relset_1(esk2_0,esk1_0,esk3_0)=k1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_12, c_0_11])).
cnf(c_0_18, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(k4_tarski(X1,X2),esk3_0)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_19, plain, (r2_hidden(esk7_2(X1,X2),X2)|r2_hidden(k4_tarski(esk7_2(X1,X2),esk8_2(X1,X2)),X1)|X2=k1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, negated_conjecture, (~r2_hidden(k4_tarski(esk4_0,X1),esk3_0)|k1_relset_1(esk2_0,esk1_0,esk3_0)!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_21, plain, (r2_hidden(k4_tarski(X1,esk6_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (X2=k1_relat_1(X1)|~r2_hidden(esk7_2(X1,X2),X2)|~r2_hidden(k4_tarski(esk7_2(X1,X2),X3),X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, negated_conjecture, (k1_relat_1(esk3_0)=esk2_0|r2_hidden(k4_tarski(X1,esk5_1(X1)),esk3_0)|~r2_hidden(X1,esk2_0)), inference(rw,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_24, negated_conjecture, (X1=k1_relat_1(esk3_0)|r2_hidden(esk7_2(esk3_0,X1),esk2_0)|r2_hidden(esk7_2(esk3_0,X1),X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25, negated_conjecture, (k1_relat_1(esk3_0)!=esk2_0|~r2_hidden(k4_tarski(esk4_0,X1),esk3_0)), inference(rw,[status(thm)],[c_0_20, c_0_17])).
cnf(c_0_26, plain, (r2_hidden(k4_tarski(X1,esk6_3(X2,k1_relat_1(X2),X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_21])).
cnf(c_0_27, negated_conjecture, (k1_relat_1(esk3_0)=esk2_0|X1=k1_relat_1(esk3_0)|~r2_hidden(esk7_2(esk3_0,X1),esk2_0)|~r2_hidden(esk7_2(esk3_0,X1),X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_28, negated_conjecture, (k1_relat_1(esk3_0)=esk2_0|r2_hidden(esk7_2(esk3_0,esk2_0),esk2_0)), inference(ef,[status(thm)],[c_0_24])).
cnf(c_0_29, negated_conjecture, (r2_hidden(esk4_0,esk2_0)|k1_relset_1(esk2_0,esk1_0,esk3_0)!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_30, negated_conjecture, (k1_relat_1(esk3_0)!=esk2_0|~r2_hidden(esk4_0,k1_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_31, negated_conjecture, (k1_relat_1(esk3_0)=esk2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_28])).
cnf(c_0_32, negated_conjecture, (r2_hidden(esk4_0,esk2_0)|k1_relat_1(esk3_0)!=esk2_0), inference(rw,[status(thm)],[c_0_29, c_0_17])).
cnf(c_0_33, negated_conjecture, (~r2_hidden(esk4_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_31]), c_0_31])])).
cnf(c_0_34, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_31])]), c_0_33]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 35
# Proof object clause steps            : 24
# Proof object formula steps           : 11
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 13
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 13
# Processed clauses                    : 89
# ...of these trivial                  : 0
# ...subsumed                          : 11
# ...remaining for further processing  : 77
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 30
# Generated clauses                    : 202
# ...of the previous two non-trivial   : 208
# Contextual simplify-reflections      : 1
# Paramodulations                      : 198
# Factorizations                       : 2
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 32
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 28
# Current number of unprocessed clauses: 145
# ...number of literals in the above   : 523
# Current number of archived formulas  : 0
# Current number of archived clauses   : 43
# Clause-clause subsumption calls (NU) : 371
# Rec. Clause-clause subsumption calls : 295
# Non-unit clause-clause subsumptions  : 11
# Unit Clause-clause subsumption calls : 5
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4417

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28644 pages
