# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', t28_xboole_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', t2_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', t48_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', t100_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', commutativity_k3_xboole_0)).
fof(t5_boole, axiom, ![X1]:k5_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', t5_boole)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', idempotence_k3_xboole_0)).
fof(t91_xboole_1, axiom, ![X1, X2, X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3)), file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', t91_xboole_1)).
fof(t13_zfmisc_1, conjecture, ![X1, X2]:(k2_xboole_0(k1_tarski(X1),k1_tarski(X2))=k1_tarski(X1)=>X1=X2), file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', t13_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', t70_enumset1)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', t98_xboole_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', t75_enumset1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', d1_tarski)).
fof(t134_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8),k1_tarski(X9)), file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', t134_enumset1)).
fof(d7_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9, X10]:(X10=k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)<=>![X11]:(r2_hidden(X11,X10)<=>~((~(X11=X1)&(~(X11=X2)&(~(X11=X3)&(~(X11=X4)&(~(X11=X5)&(~(X11=X6)&(~(X11=X7)&(~(X11=X8)&~(X11=X9)))))))))))), file('data/MPTP/proofs/6763ac03539c48a7b70c0144a462d38a.E_input', d7_enumset1)).
fof(c_0_20, plain, ![X61, X62]:(~r1_tarski(X61,X62)|k3_xboole_0(X61,X62)=X61), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_21, plain, ![X21]:r1_tarski(k1_xboole_0,X21), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_22, plain, ![X58, X59]:k4_xboole_0(X58,k4_xboole_0(X58,X59))=k3_xboole_0(X58,X59), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_23, plain, ![X63, X64]:k4_xboole_0(X63,X64)=k5_xboole_0(X63,k3_xboole_0(X63,X64)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_24, plain, ![X76, X77]:k3_xboole_0(X76,X77)=k3_xboole_0(X77,X76), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_25, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_31, plain, ![X60]:k5_xboole_0(X60,k1_xboole_0)=X60, inference(variable_rename,[status(thm)],[t5_boole])).
fof(c_0_32, plain, ![X50]:k3_xboole_0(X50,X50)=X50, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
fof(c_0_33, plain, ![X14, X15, X16]:k5_xboole_0(k5_xboole_0(X14,X15),X16)=k5_xboole_0(X14,k5_xboole_0(X15,X16)), inference(variable_rename,[status(thm)],[t91_xboole_1])).
cnf(c_0_34, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2))))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_28]), c_0_28])).
cnf(c_0_35, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_36, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_37, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_38, negated_conjecture, ~(![X1, X2]:(k2_xboole_0(k1_tarski(X1),k1_tarski(X2))=k1_tarski(X1)=>X1=X2)), inference(assume_negation,[status(cth)],[t13_zfmisc_1])).
cnf(c_0_39, plain, (k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_40, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_37])).
fof(c_0_41, negated_conjecture, (k2_xboole_0(k1_tarski(esk1_0),k1_tarski(esk2_0))=k1_tarski(esk1_0)&esk1_0!=esk2_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_38])])])).
fof(c_0_42, plain, ![X90]:k2_tarski(X90,X90)=k1_tarski(X90), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_43, plain, ![X65, X66]:k1_enumset1(X65,X65,X66)=k2_tarski(X65,X66), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_44, plain, ![X48, X49]:k2_xboole_0(X48,X49)=k5_xboole_0(X48,k4_xboole_0(X49,X48)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_45, plain, ![X45, X46, X47]:k2_enumset1(X45,X45,X46,X47)=k1_enumset1(X45,X46,X47), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_46, plain, ![X17, X18, X19, X20]:k3_enumset1(X17,X17,X18,X19,X20)=k2_enumset1(X17,X18,X19,X20), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_47, plain, ![X78, X79, X80, X81, X82]:k4_enumset1(X78,X78,X79,X80,X81,X82)=k3_enumset1(X78,X79,X80,X81,X82), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_48, plain, ![X91, X92, X93, X94, X95, X96]:k5_enumset1(X91,X91,X92,X93,X94,X95,X96)=k4_enumset1(X91,X92,X93,X94,X95,X96), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_49, plain, ![X51, X52, X53, X54, X55, X56, X57]:k6_enumset1(X51,X51,X52,X53,X54,X55,X56,X57)=k5_enumset1(X51,X52,X53,X54,X55,X56,X57), inference(variable_rename,[status(thm)],[t75_enumset1])).
cnf(c_0_50, plain, (k5_xboole_0(X1,k5_xboole_0(X1,X2))=k5_xboole_0(k1_xboole_0,X2)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_51, negated_conjecture, (k2_xboole_0(k1_tarski(esk1_0),k1_tarski(esk2_0))=k1_tarski(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_52, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_53, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_54, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_55, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_56, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_57, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_58, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_60, plain, (k5_xboole_0(k1_xboole_0,X1)=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_40]), c_0_36])).
cnf(c_0_61, negated_conjecture, (k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),k3_xboole_0(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_52]), c_0_52]), c_0_52]), c_0_53]), c_0_53]), c_0_53]), c_0_54]), c_0_28]), c_0_55]), c_0_55]), c_0_55]), c_0_55]), c_0_55]), c_0_56]), c_0_56]), c_0_56]), c_0_56]), c_0_56]), c_0_57]), c_0_57]), c_0_57]), c_0_57]), c_0_57]), c_0_58]), c_0_58]), c_0_58]), c_0_58]), c_0_58]), c_0_59]), c_0_59]), c_0_59]), c_0_59]), c_0_59])).
fof(c_0_62, plain, ![X83, X84, X85, X86, X87, X88]:(((~r2_hidden(X85,X84)|X85=X83|X84!=k1_tarski(X83))&(X86!=X83|r2_hidden(X86,X84)|X84!=k1_tarski(X83)))&((~r2_hidden(esk4_2(X87,X88),X88)|esk4_2(X87,X88)!=X87|X88=k1_tarski(X87))&(r2_hidden(esk4_2(X87,X88),X88)|esk4_2(X87,X88)=X87|X88=k1_tarski(X87)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_63, plain, ![X67, X68, X69, X70, X71, X72, X73, X74, X75]:k7_enumset1(X67,X68,X69,X70,X71,X72,X73,X74,X75)=k2_xboole_0(k6_enumset1(X67,X68,X69,X70,X71,X72,X73,X74),k1_tarski(X75)), inference(variable_rename,[status(thm)],[t134_enumset1])).
cnf(c_0_64, plain, (k5_xboole_0(X1,k5_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[c_0_50, c_0_60])).
cnf(c_0_65, negated_conjecture, (k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[c_0_61, c_0_29])).
cnf(c_0_66, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_67, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8),k1_tarski(X9))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_68, negated_conjecture, (k5_xboole_0(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_40])).
cnf(c_0_69, plain, (k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X1,X2)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_70, plain, (X1=X3|X2!=k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_52]), c_0_53]), c_0_55]), c_0_56]), c_0_57]), c_0_58]), c_0_59])).
cnf(c_0_71, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k5_xboole_0(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8),k5_xboole_0(k6_enumset1(X9,X9,X9,X9,X9,X9,X9,X9),k3_xboole_0(k6_enumset1(X9,X9,X9,X9,X9,X9,X9,X9),k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_52]), c_0_53]), c_0_54]), c_0_28]), c_0_55]), c_0_55]), c_0_56]), c_0_56]), c_0_57]), c_0_57]), c_0_58]), c_0_58]), c_0_59]), c_0_59])).
cnf(c_0_72, negated_conjecture, (k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_68]), c_0_36])).
cnf(c_0_73, plain, (k5_xboole_0(X1,k5_xboole_0(X2,X1))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_69]), c_0_36])).
fof(c_0_74, plain, ![X22, X23, X24, X25, X26, X27, X28, X29, X30, X31, X32, X33, X34, X35, X36, X37, X38, X39, X40, X41, X42, X43]:(((~r2_hidden(X32,X31)|(X32=X22|(X32=X23|(X32=X24|(X32=X25|(X32=X26|(X32=X27|(X32=X28|(X32=X29|X32=X30))))))))|X31!=k7_enumset1(X22,X23,X24,X25,X26,X27,X28,X29,X30))&((X33!=X22|r2_hidden(X33,X31)|X31!=k7_enumset1(X22,X23,X24,X25,X26,X27,X28,X29,X30))&((X33!=X23|r2_hidden(X33,X31)|X31!=k7_enumset1(X22,X23,X24,X25,X26,X27,X28,X29,X30))&((X33!=X24|r2_hidden(X33,X31)|X31!=k7_enumset1(X22,X23,X24,X25,X26,X27,X28,X29,X30))&((X33!=X25|r2_hidden(X33,X31)|X31!=k7_enumset1(X22,X23,X24,X25,X26,X27,X28,X29,X30))&((X33!=X26|r2_hidden(X33,X31)|X31!=k7_enumset1(X22,X23,X24,X25,X26,X27,X28,X29,X30))&((X33!=X27|r2_hidden(X33,X31)|X31!=k7_enumset1(X22,X23,X24,X25,X26,X27,X28,X29,X30))&((X33!=X28|r2_hidden(X33,X31)|X31!=k7_enumset1(X22,X23,X24,X25,X26,X27,X28,X29,X30))&((X33!=X29|r2_hidden(X33,X31)|X31!=k7_enumset1(X22,X23,X24,X25,X26,X27,X28,X29,X30))&(X33!=X30|r2_hidden(X33,X31)|X31!=k7_enumset1(X22,X23,X24,X25,X26,X27,X28,X29,X30)))))))))))&(((esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43)!=X34|~r2_hidden(esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43),X43)|X43=k7_enumset1(X34,X35,X36,X37,X38,X39,X40,X41,X42))&((esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43)!=X35|~r2_hidden(esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43),X43)|X43=k7_enumset1(X34,X35,X36,X37,X38,X39,X40,X41,X42))&((esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43)!=X36|~r2_hidden(esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43),X43)|X43=k7_enumset1(X34,X35,X36,X37,X38,X39,X40,X41,X42))&((esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43)!=X37|~r2_hidden(esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43),X43)|X43=k7_enumset1(X34,X35,X36,X37,X38,X39,X40,X41,X42))&((esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43)!=X38|~r2_hidden(esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43),X43)|X43=k7_enumset1(X34,X35,X36,X37,X38,X39,X40,X41,X42))&((esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43)!=X39|~r2_hidden(esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43),X43)|X43=k7_enumset1(X34,X35,X36,X37,X38,X39,X40,X41,X42))&((esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43)!=X40|~r2_hidden(esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43),X43)|X43=k7_enumset1(X34,X35,X36,X37,X38,X39,X40,X41,X42))&((esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43)!=X41|~r2_hidden(esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43),X43)|X43=k7_enumset1(X34,X35,X36,X37,X38,X39,X40,X41,X42))&(esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43)!=X42|~r2_hidden(esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43),X43)|X43=k7_enumset1(X34,X35,X36,X37,X38,X39,X40,X41,X42))))))))))&(r2_hidden(esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43),X43)|(esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43)=X34|(esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43)=X35|(esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43)=X36|(esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43)=X37|(esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43)=X38|(esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43)=X39|(esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43)=X40|(esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43)=X41|esk3_10(X34,X35,X36,X37,X38,X39,X40,X41,X42,X43)=X42))))))))|X43=k7_enumset1(X34,X35,X36,X37,X38,X39,X40,X41,X42)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_enumset1])])])])])])])).
cnf(c_0_75, plain, (X1=X2|~r2_hidden(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2))), inference(er,[status(thm)],[c_0_70])).
cnf(c_0_76, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=k7_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk1_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73])).
cnf(c_0_77, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k7_enumset1(X4,X5,X6,X7,X8,X9,X2,X10,X11)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_78, negated_conjecture, (X1=esk1_0|~r2_hidden(X1,k7_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk1_0))), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_79, plain, (r2_hidden(X1,k7_enumset1(X2,X3,X4,X5,X6,X7,X1,X8,X9))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_77])])).
cnf(c_0_80, negated_conjecture, (esk1_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_81, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 82
# Proof object clause steps            : 41
# Proof object formula steps           : 41
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 69
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 20
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 43
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 34
# Processed clauses                    : 92
# ...of these trivial                  : 2
# ...subsumed                          : 0
# ...remaining for further processing  : 90
# Other redundant clauses eliminated   : 27
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 8
# Generated clauses                    : 192
# ...of the previous two non-trivial   : 141
# Contextual simplify-reflections      : 0
# Paramodulations                      : 175
# Factorizations                       : 0
# Equation resolutions                 : 27
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 36
#    Positive orientable unit clauses  : 24
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 10
# Current number of unprocessed clauses: 112
# ...number of literals in the above   : 302
# Current number of archived formulas  : 0
# Current number of archived clauses   : 51
# Clause-clause subsumption calls (NU) : 84
# Rec. Clause-clause subsumption calls : 84
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 101
# BW rewrite match successes           : 10
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4831

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28656 pages
