# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2)))))))), file('data/MPTP/proofs/679eb3b4b59e41329052b5ff469f9e68.E_input', t7_yellow_0)).
fof(t30_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4)))))&((r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))=>(X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3)))))), file('data/MPTP/proofs/679eb3b4b59e41329052b5ff469f9e68.E_input', t30_yellow_0)).
fof(t39_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2))), file('data/MPTP/proofs/679eb3b4b59e41329052b5ff469f9e68.E_input', t39_yellow_0)).
fof(t31_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3))=>(r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2)))))&((r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2))))=>(X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3)))))), file('data/MPTP/proofs/679eb3b4b59e41329052b5ff469f9e68.E_input', t31_yellow_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/679eb3b4b59e41329052b5ff469f9e68.E_input', t24_orders_2)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/679eb3b4b59e41329052b5ff469f9e68.E_input', t5_subset)).
fof(dt_u1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))), file('data/MPTP/proofs/679eb3b4b59e41329052b5ff469f9e68.E_input', dt_u1_orders_2)).
fof(cc3_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/679eb3b4b59e41329052b5ff469f9e68.E_input', cc3_relset_1)).
fof(rc1_subset_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>?[X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))&~(v1_xboole_0(X2)))), file('data/MPTP/proofs/679eb3b4b59e41329052b5ff469f9e68.E_input', rc1_subset_1)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/679eb3b4b59e41329052b5ff469f9e68.E_input', redefinition_k6_domain_1)).
fof(d9_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)<=>r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1)))))), file('data/MPTP/proofs/679eb3b4b59e41329052b5ff469f9e68.E_input', d9_orders_2)).
fof(c_0_11, plain, ![X21, X22, X23]:((~r1_lattice3(X21,k1_tarski(X23),X22)|r1_orders_2(X21,X22,X23)|~m1_subset_1(X23,u1_struct_0(X21))|~m1_subset_1(X22,u1_struct_0(X21))|~l1_orders_2(X21))&((~r1_orders_2(X21,X22,X23)|r1_lattice3(X21,k1_tarski(X23),X22)|~m1_subset_1(X23,u1_struct_0(X21))|~m1_subset_1(X22,u1_struct_0(X21))|~l1_orders_2(X21))&((~r2_lattice3(X21,k1_tarski(X23),X22)|r1_orders_2(X21,X23,X22)|~m1_subset_1(X23,u1_struct_0(X21))|~m1_subset_1(X22,u1_struct_0(X21))|~l1_orders_2(X21))&(~r1_orders_2(X21,X23,X22)|r2_lattice3(X21,k1_tarski(X23),X22)|~m1_subset_1(X23,u1_struct_0(X21))|~m1_subset_1(X22,u1_struct_0(X21))|~l1_orders_2(X21))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])).
fof(c_0_12, plain, ![X7, X8, X9, X10, X11]:(((r2_lattice3(X7,X9,X8)|(X8!=k1_yellow_0(X7,X9)|~r1_yellow_0(X7,X9))|~m1_subset_1(X8,u1_struct_0(X7))|(~v5_orders_2(X7)|~l1_orders_2(X7)))&(~m1_subset_1(X10,u1_struct_0(X7))|(~r2_lattice3(X7,X9,X10)|r1_orders_2(X7,X8,X10))|(X8!=k1_yellow_0(X7,X9)|~r1_yellow_0(X7,X9))|~m1_subset_1(X8,u1_struct_0(X7))|(~v5_orders_2(X7)|~l1_orders_2(X7))))&(((X8=k1_yellow_0(X7,X11)|(m1_subset_1(esk3_3(X7,X8,X11),u1_struct_0(X7))|~r2_lattice3(X7,X11,X8))|~m1_subset_1(X8,u1_struct_0(X7))|(~v5_orders_2(X7)|~l1_orders_2(X7)))&(r1_yellow_0(X7,X11)|(m1_subset_1(esk3_3(X7,X8,X11),u1_struct_0(X7))|~r2_lattice3(X7,X11,X8))|~m1_subset_1(X8,u1_struct_0(X7))|(~v5_orders_2(X7)|~l1_orders_2(X7))))&(((X8=k1_yellow_0(X7,X11)|(r2_lattice3(X7,X11,esk3_3(X7,X8,X11))|~r2_lattice3(X7,X11,X8))|~m1_subset_1(X8,u1_struct_0(X7))|(~v5_orders_2(X7)|~l1_orders_2(X7)))&(r1_yellow_0(X7,X11)|(r2_lattice3(X7,X11,esk3_3(X7,X8,X11))|~r2_lattice3(X7,X11,X8))|~m1_subset_1(X8,u1_struct_0(X7))|(~v5_orders_2(X7)|~l1_orders_2(X7))))&((X8=k1_yellow_0(X7,X11)|(~r1_orders_2(X7,X8,esk3_3(X7,X8,X11))|~r2_lattice3(X7,X11,X8))|~m1_subset_1(X8,u1_struct_0(X7))|(~v5_orders_2(X7)|~l1_orders_2(X7)))&(r1_yellow_0(X7,X11)|(~r1_orders_2(X7,X8,esk3_3(X7,X8,X11))|~r2_lattice3(X7,X11,X8))|~m1_subset_1(X8,u1_struct_0(X7))|(~v5_orders_2(X7)|~l1_orders_2(X7))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_yellow_0])])])])])])).
cnf(c_0_13, plain, (r1_orders_2(X1,X2,X3)|~r2_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_14, plain, (X1=k1_yellow_0(X2,X3)|r2_lattice3(X2,X3,esk3_3(X2,X1,X3))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15, plain, (X1=k1_yellow_0(X2,X3)|m1_subset_1(esk3_3(X2,X1,X3),u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_16, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2)))), inference(assume_negation,[status(cth)],[t39_yellow_0])).
fof(c_0_17, plain, ![X13, X14, X15, X16, X17]:(((r1_lattice3(X13,X15,X14)|(X14!=k2_yellow_0(X13,X15)|~r2_yellow_0(X13,X15))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13)))&(~m1_subset_1(X16,u1_struct_0(X13))|(~r1_lattice3(X13,X15,X16)|r1_orders_2(X13,X16,X14))|(X14!=k2_yellow_0(X13,X15)|~r2_yellow_0(X13,X15))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13))))&(((X14=k2_yellow_0(X13,X17)|(m1_subset_1(esk4_3(X13,X14,X17),u1_struct_0(X13))|~r1_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13)))&(r2_yellow_0(X13,X17)|(m1_subset_1(esk4_3(X13,X14,X17),u1_struct_0(X13))|~r1_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13))))&(((X14=k2_yellow_0(X13,X17)|(r1_lattice3(X13,X17,esk4_3(X13,X14,X17))|~r1_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13)))&(r2_yellow_0(X13,X17)|(r1_lattice3(X13,X17,esk4_3(X13,X14,X17))|~r1_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13))))&((X14=k2_yellow_0(X13,X17)|(~r1_orders_2(X13,esk4_3(X13,X14,X17),X14)|~r1_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13)))&(r2_yellow_0(X13,X17)|(~r1_orders_2(X13,esk4_3(X13,X14,X17),X14)|~r1_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_yellow_0])])])])])])).
cnf(c_0_18, plain, (X1=k1_yellow_0(X2,X3)|~r1_orders_2(X2,X1,esk3_3(X2,X1,X3))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, plain, (X1=k1_yellow_0(X2,k1_tarski(X3))|r1_orders_2(X2,X3,esk3_3(X2,X1,k1_tarski(X3)))|~r2_lattice3(X2,k1_tarski(X3),X1)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])).
fof(c_0_20, plain, ![X34, X35]:(v2_struct_0(X34)|(~v3_orders_2(X34)|~l1_orders_2(X34))|(~m1_subset_1(X35,u1_struct_0(X34))|r1_orders_2(X34,X35,X35))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_21, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_16])])])])).
cnf(c_0_22, plain, (r1_orders_2(X1,X3,X2)|~r1_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_23, plain, (X1=k2_yellow_0(X2,X3)|r1_lattice3(X2,X3,esk4_3(X2,X1,X3))|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (X1=k2_yellow_0(X2,X3)|m1_subset_1(esk4_3(X2,X1,X3),u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_25, plain, ![X28, X29, X30]:(~r2_hidden(X28,X29)|(~m1_subset_1(X29,k1_zfmisc_1(X30))|~v1_xboole_0(X30))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_26, plain, ![X24]:(~l1_orders_2(X24)|m1_subset_1(u1_orders_2(X24),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X24),u1_struct_0(X24))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_orders_2])])).
fof(c_0_27, plain, ![X31, X32, X33]:(~v1_xboole_0(X31)|(~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X32)))|v1_xboole_0(X33))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_relset_1])])])).
fof(c_0_28, plain, ![X19]:((m1_subset_1(esk5_1(X19),k1_zfmisc_1(X19))|v1_xboole_0(X19))&(~v1_xboole_0(esk5_1(X19))|v1_xboole_0(X19))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[rc1_subset_1])])])])])).
fof(c_0_29, plain, ![X36, X37]:(v1_xboole_0(X36)|~m1_subset_1(X37,X36)|k6_domain_1(X36,X37)=k1_tarski(X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_30, plain, (k1_yellow_0(X1,k1_tarski(X2))=X2|~r2_lattice3(X1,k1_tarski(X2),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_31, plain, (r2_lattice3(X1,k1_tarski(X2),X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_32, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_33, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_34, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_35, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_36, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_37, plain, (X1=k2_yellow_0(X2,X3)|~r1_orders_2(X2,esk4_3(X2,X1,X3),X1)|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_38, plain, (X1=k2_yellow_0(X2,k1_tarski(X3))|r1_orders_2(X2,esk4_3(X2,X1,k1_tarski(X3)),X3)|~r1_lattice3(X2,k1_tarski(X3),X1)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24])).
cnf(c_0_39, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_40, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_41, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_42, plain, (m1_subset_1(esk5_1(X1),k1_zfmisc_1(X1))|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_43, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk5_1(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_44, negated_conjecture, (k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_45, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_46, plain, (k1_yellow_0(X1,k1_tarski(X2))=X2|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_47, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34]), c_0_35])]), c_0_36])).
cnf(c_0_48, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_49, plain, (k2_yellow_0(X1,k1_tarski(X2))=X2|~r1_lattice3(X1,k1_tarski(X2),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_50, plain, (r1_lattice3(X1,k1_tarski(X3),X2)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_51, plain, (~r2_hidden(X1,u1_orders_2(X2))|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_52, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43])).
fof(c_0_53, plain, ![X25, X26, X27]:((~r1_orders_2(X25,X26,X27)|r2_hidden(k4_tarski(X26,X27),u1_orders_2(X25))|~m1_subset_1(X27,u1_struct_0(X25))|~m1_subset_1(X26,u1_struct_0(X25))|~l1_orders_2(X25))&(~r2_hidden(k4_tarski(X26,X27),u1_orders_2(X25))|r1_orders_2(X25,X26,X27)|~m1_subset_1(X27,u1_struct_0(X25))|~m1_subset_1(X26,u1_struct_0(X25))|~l1_orders_2(X25))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_orders_2])])])])).
cnf(c_0_54, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|k1_yellow_0(esk1_0,k1_tarski(esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k1_tarski(esk2_0))!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_33])])).
cnf(c_0_55, negated_conjecture, (k1_yellow_0(esk1_0,k1_tarski(esk2_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_33]), c_0_34]), c_0_48])])).
cnf(c_0_56, plain, (k2_yellow_0(X1,k1_tarski(X2))=X2|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_57, plain, (~r2_hidden(X1,u1_orders_2(X2))|~v1_xboole_0(u1_struct_0(X2))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_58, plain, (r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_59, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|k2_yellow_0(esk1_0,k1_tarski(esk2_0))!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_55])])).
cnf(c_0_60, negated_conjecture, (k2_yellow_0(esk1_0,k1_tarski(esk2_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_47]), c_0_33]), c_0_34]), c_0_48])])).
cnf(c_0_61, plain, (~v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_62, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_60])])).
cnf(c_0_63, negated_conjecture, (~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_34])])).
cnf(c_0_64, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_47]), c_0_33])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 65
# Proof object clause steps            : 42
# Proof object formula steps           : 23
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 25
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 35
# Processed clauses                    : 123
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 121
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 124
# ...of the previous two non-trivial   : 95
# Contextual simplify-reflections      : 12
# Paramodulations                      : 120
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 80
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 67
# Current number of unprocessed clauses: 30
# ...number of literals in the above   : 231
# Current number of archived formulas  : 0
# Current number of archived clauses   : 37
# Clause-clause subsumption calls (NU) : 2814
# Rec. Clause-clause subsumption calls : 147
# Non-unit clause-clause subsumptions  : 14
# Unit Clause-clause subsumption calls : 14
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7400

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.003 s
# Total time               : 0.014 s
# Maximum resident set size: 28816 pages
