# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t71_yellow_0, conjecture, ![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&(v6_yellow_0(X2,X1)&m1_yellow_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>![X6]:(m1_subset_1(X6,u1_struct_0(X1))=>((X5=X3&X6=X4)=>k13_lattice3(X2,X3,X4)=k13_lattice3(X1,X5,X6)))))))), file('data/MPTP/proofs/67bf3b18c1a94d3fa60fc82cf2d56c73.E_input', t71_yellow_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/67bf3b18c1a94d3fa60fc82cf2d56c73.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/67bf3b18c1a94d3fa60fc82cf2d56c73.E_input', dt_l1_orders_2)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/67bf3b18c1a94d3fa60fc82cf2d56c73.E_input', t2_subset)).
fof(dt_m1_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_yellow_0(X2,X1)=>l1_orders_2(X2))), file('data/MPTP/proofs/67bf3b18c1a94d3fa60fc82cf2d56c73.E_input', dt_m1_yellow_0)).
fof(cc1_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v1_lattice3(X1)=>~(v2_struct_0(X1)))), file('data/MPTP/proofs/67bf3b18c1a94d3fa60fc82cf2d56c73.E_input', cc1_lattice3)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/67bf3b18c1a94d3fa60fc82cf2d56c73.E_input', d1_xboole_0)).
fof(redefinition_k7_domain_1, axiom, ![X1, X2, X3]:((~(v1_xboole_0(X1))&(m1_subset_1(X2,X1)&m1_subset_1(X3,X1)))=>k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)), file('data/MPTP/proofs/67bf3b18c1a94d3fa60fc82cf2d56c73.E_input', redefinition_k7_domain_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/67bf3b18c1a94d3fa60fc82cf2d56c73.E_input', commutativity_k2_tarski)).
fof(cc12_yellow_0, axiom, ![X1]:((v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_yellow_0(X2,X1)=>((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&v6_yellow_0(X2,X1)))=>(~(v2_struct_0(X2))&(v1_lattice3(X2)&(v4_yellow_0(X2,X1)&v6_yellow_0(X2,X1))))))), file('data/MPTP/proofs/67bf3b18c1a94d3fa60fc82cf2d56c73.E_input', cc12_yellow_0)).
fof(cc8_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_yellow_0(X2,X1)=>(v4_yellow_0(X2,X1)=>(v5_orders_2(X2)&v4_yellow_0(X2,X1))))), file('data/MPTP/proofs/67bf3b18c1a94d3fa60fc82cf2d56c73.E_input', cc8_yellow_0)).
fof(cc7_yellow_0, axiom, ![X1]:((v4_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_yellow_0(X2,X1)=>(v4_yellow_0(X2,X1)=>(v4_orders_2(X2)&v4_yellow_0(X2,X1))))), file('data/MPTP/proofs/67bf3b18c1a94d3fa60fc82cf2d56c73.E_input', cc7_yellow_0)).
fof(cc6_yellow_0, axiom, ![X1]:((v3_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_yellow_0(X2,X1)=>(v4_yellow_0(X2,X1)=>(v3_orders_2(X2)&v4_yellow_0(X2,X1))))), file('data/MPTP/proofs/67bf3b18c1a94d3fa60fc82cf2d56c73.E_input', cc6_yellow_0)).
fof(t41_yellow_0, axiom, ![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k1_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k13_lattice3(X1,X2,X3)))), file('data/MPTP/proofs/67bf3b18c1a94d3fa60fc82cf2d56c73.E_input', t41_yellow_0)).
fof(t67_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>((r1_yellow_0(X1,k7_domain_1(u1_struct_0(X2),X3,X4))&r2_hidden(k1_yellow_0(X1,k7_domain_1(u1_struct_0(X2),X3,X4)),u1_struct_0(X2)))=>(r1_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))&k1_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))=k1_yellow_0(X1,k7_domain_1(u1_struct_0(X2),X3,X4)))))))), file('data/MPTP/proofs/67bf3b18c1a94d3fa60fc82cf2d56c73.E_input', t67_yellow_0)).
fof(d17_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_yellow_0(X2,X1)=>(v6_yellow_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,u1_struct_0(X2))&(r2_hidden(X4,u1_struct_0(X2))&r1_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))))=>r2_hidden(k1_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4)),u1_struct_0(X2)))))))), file('data/MPTP/proofs/67bf3b18c1a94d3fa60fc82cf2d56c73.E_input', d17_yellow_0)).
fof(t20_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>(v1_lattice3(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>r1_yellow_0(X1,k2_tarski(X2,X3)))))), file('data/MPTP/proofs/67bf3b18c1a94d3fa60fc82cf2d56c73.E_input', t20_yellow_0)).
fof(c_0_17, negated_conjecture, ~(![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&(v6_yellow_0(X2,X1)&m1_yellow_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>![X6]:(m1_subset_1(X6,u1_struct_0(X1))=>((X5=X3&X6=X4)=>k13_lattice3(X2,X3,X4)=k13_lattice3(X1,X5,X6))))))))), inference(assume_negation,[status(cth)],[t71_yellow_0])).
fof(c_0_18, negated_conjecture, ((v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_lattice3(esk1_0)&l1_orders_2(esk1_0)))))&((~v2_struct_0(esk2_0)&(v4_yellow_0(esk2_0,esk1_0)&(v6_yellow_0(esk2_0,esk1_0)&m1_yellow_0(esk2_0,esk1_0))))&(m1_subset_1(esk3_0,u1_struct_0(esk2_0))&(m1_subset_1(esk4_0,u1_struct_0(esk2_0))&(m1_subset_1(esk5_0,u1_struct_0(esk1_0))&(m1_subset_1(esk6_0,u1_struct_0(esk1_0))&((esk5_0=esk3_0&esk6_0=esk4_0)&k13_lattice3(esk2_0,esk3_0,esk4_0)!=k13_lattice3(esk1_0,esk5_0,esk6_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])).
fof(c_0_19, plain, ![X40]:(v2_struct_0(X40)|~l1_struct_0(X40)|~v1_xboole_0(u1_struct_0(X40))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_20, plain, ![X39]:(~l1_orders_2(X39)|l1_struct_0(X39)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
fof(c_0_21, plain, ![X41, X42]:(~m1_subset_1(X41,X42)|(v1_xboole_0(X42)|r2_hidden(X41,X42))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_22, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, negated_conjecture, (esk6_0=esk4_0), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_24, plain, ![X53, X54]:(~l1_orders_2(X53)|(~m1_yellow_0(X54,X53)|l1_orders_2(X54))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_yellow_0])])])).
fof(c_0_25, plain, ![X19]:(~l1_orders_2(X19)|(~v1_lattice3(X19)|~v2_struct_0(X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattice3])])])).
cnf(c_0_26, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk6_0,u1_struct_0(esk2_0))), inference(rw,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_30, plain, (l1_orders_2(X2)|~l1_orders_2(X1)|~m1_yellow_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, negated_conjecture, (m1_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_32, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_33, negated_conjecture, (m1_subset_1(esk6_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_34, plain, (~l1_orders_2(X1)|~v1_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, negated_conjecture, (v1_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_36, plain, ![X27, X28, X29]:((~v1_xboole_0(X27)|~r2_hidden(X28,X27))&(r2_hidden(esk11_1(X29),X29)|v1_xboole_0(X29))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_37, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_38, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|r2_hidden(esk6_0,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_39, negated_conjecture, (l1_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])])).
cnf(c_0_40, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_41, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk6_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_28, c_0_33])).
cnf(c_0_42, negated_conjecture, (~v2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_32])])).
fof(c_0_43, plain, ![X43, X44, X45]:(v1_xboole_0(X43)|(~m1_subset_1(X44,X43)|~m1_subset_1(X45,X43))|k7_domain_1(X43,X44,X45)=k2_tarski(X44,X45)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k7_domain_1])])])).
cnf(c_0_44, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_45, negated_conjecture, (r2_hidden(esk6_0,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])]), c_0_40])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk6_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_41]), c_0_32])]), c_0_42])).
fof(c_0_47, plain, ![X31, X32]:k2_tarski(X31,X32)=k2_tarski(X32,X31), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_48, plain, (v1_xboole_0(X1)|k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_50, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_51, negated_conjecture, (esk5_0=esk3_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_52, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_44, c_0_46])).
cnf(c_0_53, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_54, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_55, negated_conjecture, (k2_tarski(X1,esk6_0)=k7_domain_1(u1_struct_0(esk2_0),X1,esk6_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_29]), c_0_49])).
cnf(c_0_56, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(rw,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_57, negated_conjecture, (k2_tarski(X1,esk6_0)=k7_domain_1(u1_struct_0(esk1_0),X1,esk6_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_33]), c_0_52])).
cnf(c_0_58, negated_conjecture, (k2_tarski(X1,esk3_0)=k7_domain_1(u1_struct_0(esk2_0),X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_53]), c_0_49])).
cnf(c_0_59, negated_conjecture, (k2_tarski(esk6_0,X1)=k7_domain_1(u1_struct_0(esk2_0),X1,esk6_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_60, negated_conjecture, (k2_tarski(X1,esk3_0)=k7_domain_1(u1_struct_0(esk1_0),X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_56]), c_0_52])).
cnf(c_0_61, negated_conjecture, (k2_tarski(esk6_0,X1)=k7_domain_1(u1_struct_0(esk1_0),X1,esk6_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_54, c_0_57])).
cnf(c_0_62, negated_conjecture, (k7_domain_1(u1_struct_0(esk2_0),esk6_0,esk3_0)=k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_29]), c_0_53])])).
cnf(c_0_63, negated_conjecture, (k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0)=k7_domain_1(u1_struct_0(esk1_0),esk6_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_59]), c_0_33]), c_0_53])])).
fof(c_0_64, plain, ![X25, X26]:((~v2_struct_0(X26)|(v2_struct_0(X26)|(~v4_yellow_0(X26,X25)|~v6_yellow_0(X26,X25)))|~m1_yellow_0(X26,X25)|(~v4_orders_2(X25)|(~v5_orders_2(X25)|(~v1_lattice3(X25)|~l1_orders_2(X25)))))&((v1_lattice3(X26)|(v2_struct_0(X26)|(~v4_yellow_0(X26,X25)|~v6_yellow_0(X26,X25)))|~m1_yellow_0(X26,X25)|(~v4_orders_2(X25)|(~v5_orders_2(X25)|(~v1_lattice3(X25)|~l1_orders_2(X25)))))&((v4_yellow_0(X26,X25)|(v2_struct_0(X26)|(~v4_yellow_0(X26,X25)|~v6_yellow_0(X26,X25)))|~m1_yellow_0(X26,X25)|(~v4_orders_2(X25)|(~v5_orders_2(X25)|(~v1_lattice3(X25)|~l1_orders_2(X25)))))&(v6_yellow_0(X26,X25)|(v2_struct_0(X26)|(~v4_yellow_0(X26,X25)|~v6_yellow_0(X26,X25)))|~m1_yellow_0(X26,X25)|(~v4_orders_2(X25)|(~v5_orders_2(X25)|(~v1_lattice3(X25)|~l1_orders_2(X25)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc12_yellow_0])])])])])).
fof(c_0_65, plain, ![X37, X38]:((v5_orders_2(X38)|~v4_yellow_0(X38,X37)|~m1_yellow_0(X38,X37)|(~v5_orders_2(X37)|~l1_orders_2(X37)))&(v4_yellow_0(X38,X37)|~v4_yellow_0(X38,X37)|~m1_yellow_0(X38,X37)|(~v5_orders_2(X37)|~l1_orders_2(X37)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_yellow_0])])])])).
fof(c_0_66, plain, ![X51, X52]:((v4_orders_2(X52)|~v4_yellow_0(X52,X51)|~m1_yellow_0(X52,X51)|(~v4_orders_2(X51)|~l1_orders_2(X51)))&(v4_yellow_0(X52,X51)|~v4_yellow_0(X52,X51)|~m1_yellow_0(X52,X51)|(~v4_orders_2(X51)|~l1_orders_2(X51)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc7_yellow_0])])])])).
fof(c_0_67, plain, ![X49, X50]:((v3_orders_2(X50)|~v4_yellow_0(X50,X49)|~m1_yellow_0(X50,X49)|(~v3_orders_2(X49)|~l1_orders_2(X49)))&(v4_yellow_0(X50,X49)|~v4_yellow_0(X50,X49)|~m1_yellow_0(X50,X49)|(~v3_orders_2(X49)|~l1_orders_2(X49)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc6_yellow_0])])])])).
fof(c_0_68, plain, ![X46, X47, X48]:(~v3_orders_2(X46)|(~v4_orders_2(X46)|(~v5_orders_2(X46)|(~v1_lattice3(X46)|~l1_orders_2(X46))))|(~m1_subset_1(X47,u1_struct_0(X46))|(~m1_subset_1(X48,u1_struct_0(X46))|k1_yellow_0(X46,k7_domain_1(u1_struct_0(X46),X47,X48))=k13_lattice3(X46,X47,X48)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t41_yellow_0])])])).
fof(c_0_69, plain, ![X33, X34, X35, X36]:((r1_yellow_0(X34,k7_domain_1(u1_struct_0(X34),X35,X36))|(~r1_yellow_0(X33,k7_domain_1(u1_struct_0(X34),X35,X36))|~r2_hidden(k1_yellow_0(X33,k7_domain_1(u1_struct_0(X34),X35,X36)),u1_struct_0(X34)))|~m1_subset_1(X36,u1_struct_0(X34))|~m1_subset_1(X35,u1_struct_0(X34))|(v2_struct_0(X34)|(~v4_yellow_0(X34,X33)|~m1_yellow_0(X34,X33)))|(v2_struct_0(X33)|(~v4_orders_2(X33)|~l1_orders_2(X33))))&(k1_yellow_0(X34,k7_domain_1(u1_struct_0(X34),X35,X36))=k1_yellow_0(X33,k7_domain_1(u1_struct_0(X34),X35,X36))|(~r1_yellow_0(X33,k7_domain_1(u1_struct_0(X34),X35,X36))|~r2_hidden(k1_yellow_0(X33,k7_domain_1(u1_struct_0(X34),X35,X36)),u1_struct_0(X34)))|~m1_subset_1(X36,u1_struct_0(X34))|~m1_subset_1(X35,u1_struct_0(X34))|(v2_struct_0(X34)|(~v4_yellow_0(X34,X33)|~m1_yellow_0(X34,X33)))|(v2_struct_0(X33)|(~v4_orders_2(X33)|~l1_orders_2(X33))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t67_yellow_0])])])])])).
cnf(c_0_70, negated_conjecture, (k7_domain_1(u1_struct_0(esk1_0),esk6_0,esk3_0)=k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_61]), c_0_62]), c_0_29]), c_0_56])]), c_0_63])).
cnf(c_0_71, plain, (v1_lattice3(X1)|v2_struct_0(X1)|~v4_yellow_0(X1,X2)|~v6_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_72, negated_conjecture, (v6_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_73, negated_conjecture, (v4_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_74, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_75, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_76, plain, (v5_orders_2(X1)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_77, plain, (v4_orders_2(X1)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~v4_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_78, plain, (v3_orders_2(X1)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~v3_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_79, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_80, plain, (k1_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k13_lattice3(X1,X2,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_81, plain, (k1_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k1_yellow_0(X4,k7_domain_1(u1_struct_0(X1),X2,X3))|v2_struct_0(X1)|v2_struct_0(X4)|~r1_yellow_0(X4,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_hidden(k1_yellow_0(X4,k7_domain_1(u1_struct_0(X1),X2,X3)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_yellow_0(X1,X4)|~m1_yellow_0(X1,X4)|~v4_orders_2(X4)|~l1_orders_2(X4)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_82, negated_conjecture, (k7_domain_1(u1_struct_0(esk2_0),esk6_0,esk3_0)=k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_63]), c_0_70])).
cnf(c_0_83, negated_conjecture, (v1_lattice3(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_31]), c_0_72]), c_0_73]), c_0_32]), c_0_35]), c_0_74]), c_0_75])]), c_0_40])).
cnf(c_0_84, negated_conjecture, (v5_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_73]), c_0_31]), c_0_32]), c_0_74])])).
cnf(c_0_85, negated_conjecture, (v4_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_73]), c_0_31]), c_0_32]), c_0_75])])).
cnf(c_0_86, negated_conjecture, (v3_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_31]), c_0_73]), c_0_32]), c_0_79])])).
cnf(c_0_87, negated_conjecture, (k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))=k13_lattice3(esk1_0,esk6_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_70]), c_0_56]), c_0_33]), c_0_32]), c_0_35]), c_0_74]), c_0_75]), c_0_79])])).
fof(c_0_88, plain, ![X13, X14, X15, X16]:((~v6_yellow_0(X14,X13)|(~m1_subset_1(X15,u1_struct_0(X13))|(~m1_subset_1(X16,u1_struct_0(X13))|(~r2_hidden(X15,u1_struct_0(X14))|(~r2_hidden(X16,u1_struct_0(X14))|~r1_yellow_0(X13,k7_domain_1(u1_struct_0(X13),X15,X16)))|r2_hidden(k1_yellow_0(X13,k7_domain_1(u1_struct_0(X13),X15,X16)),u1_struct_0(X14)))))|~m1_yellow_0(X14,X13)|(v2_struct_0(X13)|~l1_orders_2(X13)))&((m1_subset_1(esk7_2(X13,X14),u1_struct_0(X13))|v6_yellow_0(X14,X13)|~m1_yellow_0(X14,X13)|(v2_struct_0(X13)|~l1_orders_2(X13)))&((m1_subset_1(esk8_2(X13,X14),u1_struct_0(X13))|v6_yellow_0(X14,X13)|~m1_yellow_0(X14,X13)|(v2_struct_0(X13)|~l1_orders_2(X13)))&(((r2_hidden(esk7_2(X13,X14),u1_struct_0(X14))|v6_yellow_0(X14,X13)|~m1_yellow_0(X14,X13)|(v2_struct_0(X13)|~l1_orders_2(X13)))&((r2_hidden(esk8_2(X13,X14),u1_struct_0(X14))|v6_yellow_0(X14,X13)|~m1_yellow_0(X14,X13)|(v2_struct_0(X13)|~l1_orders_2(X13)))&(r1_yellow_0(X13,k7_domain_1(u1_struct_0(X13),esk7_2(X13,X14),esk8_2(X13,X14)))|v6_yellow_0(X14,X13)|~m1_yellow_0(X14,X13)|(v2_struct_0(X13)|~l1_orders_2(X13)))))&(~r2_hidden(k1_yellow_0(X13,k7_domain_1(u1_struct_0(X13),esk7_2(X13,X14),esk8_2(X13,X14))),u1_struct_0(X14))|v6_yellow_0(X14,X13)|~m1_yellow_0(X14,X13)|(v2_struct_0(X13)|~l1_orders_2(X13))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_yellow_0])])])])])])).
cnf(c_0_89, negated_conjecture, (k1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),X1,esk6_0))=k1_yellow_0(X2,k7_domain_1(u1_struct_0(esk2_0),X1,esk6_0))|v2_struct_0(X2)|~r1_yellow_0(X2,k7_domain_1(u1_struct_0(esk2_0),X1,esk6_0))|~r2_hidden(k1_yellow_0(X2,k7_domain_1(u1_struct_0(esk2_0),X1,esk6_0)),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_yellow_0(esk2_0,X2)|~v4_yellow_0(esk2_0,X2)|~l1_orders_2(X2)|~v4_orders_2(X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_29]), c_0_40])).
cnf(c_0_90, negated_conjecture, (k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0)=k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)), inference(rw,[status(thm)],[c_0_63, c_0_70])).
cnf(c_0_91, negated_conjecture, (k1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))=k13_lattice3(esk2_0,esk6_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_82]), c_0_53]), c_0_29]), c_0_39]), c_0_83]), c_0_84]), c_0_85]), c_0_86])])).
cnf(c_0_92, negated_conjecture, (k13_lattice3(esk1_0,esk6_0,esk3_0)=k13_lattice3(esk1_0,esk3_0,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_87]), c_0_33]), c_0_56]), c_0_32]), c_0_35]), c_0_74]), c_0_75]), c_0_79])])).
cnf(c_0_93, negated_conjecture, (k13_lattice3(esk2_0,esk3_0,esk4_0)!=k13_lattice3(esk1_0,esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_94, plain, (r2_hidden(k1_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4)),u1_struct_0(X1))|v2_struct_0(X2)|~v6_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,u1_struct_0(X1))|~r2_hidden(X4,u1_struct_0(X1))|~r1_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))|~m1_yellow_0(X1,X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_95, negated_conjecture, (k1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),X1,esk6_0))=k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),X1,esk6_0))|~r1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),X1,esk6_0))|~r2_hidden(k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),X1,esk6_0)),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_31]), c_0_73]), c_0_32]), c_0_75])]), c_0_42])).
cnf(c_0_96, negated_conjecture, (k13_lattice3(esk2_0,esk6_0,esk3_0)=k13_lattice3(esk2_0,esk3_0,esk6_0)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_90]), c_0_29]), c_0_53]), c_0_39]), c_0_83]), c_0_84]), c_0_85]), c_0_86])]), c_0_91])).
cnf(c_0_97, negated_conjecture, (k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))=k13_lattice3(esk1_0,esk3_0,esk6_0)), inference(rw,[status(thm)],[c_0_87, c_0_92])).
cnf(c_0_98, negated_conjecture, (k13_lattice3(esk2_0,esk3_0,esk6_0)!=k13_lattice3(esk1_0,esk3_0,esk6_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_23]), c_0_51])).
cnf(c_0_99, negated_conjecture, (r2_hidden(k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,esk6_0)),u1_struct_0(X2))|~r1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,esk6_0))|~r2_hidden(esk6_0,u1_struct_0(X2))|~r2_hidden(X1,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_yellow_0(X2,esk1_0)|~v6_yellow_0(X2,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_33]), c_0_32])]), c_0_42])).
fof(c_0_100, plain, ![X20, X21, X22]:((~v1_lattice3(X20)|(~m1_subset_1(X21,u1_struct_0(X20))|(~m1_subset_1(X22,u1_struct_0(X20))|r1_yellow_0(X20,k2_tarski(X21,X22))))|(~v5_orders_2(X20)|~l1_orders_2(X20)))&((m1_subset_1(esk9_1(X20),u1_struct_0(X20))|v1_lattice3(X20)|(~v5_orders_2(X20)|~l1_orders_2(X20)))&((m1_subset_1(esk10_1(X20),u1_struct_0(X20))|v1_lattice3(X20)|(~v5_orders_2(X20)|~l1_orders_2(X20)))&(~r1_yellow_0(X20,k2_tarski(esk9_1(X20),esk10_1(X20)))|v1_lattice3(X20)|(~v5_orders_2(X20)|~l1_orders_2(X20)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t20_yellow_0])])])])])).
cnf(c_0_101, negated_conjecture, (~r1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))|~r2_hidden(k13_lattice3(esk1_0,esk3_0,esk6_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_90]), c_0_91]), c_0_96]), c_0_97]), c_0_97]), c_0_53])]), c_0_98])).
cnf(c_0_102, negated_conjecture, (r2_hidden(k13_lattice3(esk1_0,X1,esk6_0),u1_struct_0(X2))|~r1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,esk6_0))|~r2_hidden(esk6_0,u1_struct_0(X2))|~r2_hidden(X1,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_yellow_0(X2,esk1_0)|~v6_yellow_0(X2,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_80]), c_0_33]), c_0_32]), c_0_35]), c_0_74]), c_0_75]), c_0_79])])).
cnf(c_0_103, negated_conjecture, (r2_hidden(esk3_0,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_53]), c_0_49])).
cnf(c_0_104, plain, (r1_yellow_0(X1,k2_tarski(X2,X3))|~v1_lattice3(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_105, negated_conjecture, (~r1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_45]), c_0_103]), c_0_56]), c_0_31]), c_0_72])])).
cnf(c_0_106, negated_conjecture, (r1_yellow_0(X1,k7_domain_1(u1_struct_0(esk1_0),X2,esk6_0))|~m1_subset_1(esk6_0,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_104, c_0_57])).
cnf(c_0_107, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_33]), c_0_56]), c_0_32]), c_0_35]), c_0_74])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 108
# Proof object clause steps            : 73
# Proof object formula steps           : 35
# Proof object conjectures             : 59
# Proof object clause conjectures      : 56
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 32
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 35
# Proof object simplifying inferences  : 122
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 49
# Removed in clause preprocessing      : 6
# Initial clauses in saturation        : 43
# Processed clauses                    : 246
# ...of these trivial                  : 2
# ...subsumed                          : 59
# ...remaining for further processing  : 185
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 6
# Generated clauses                    : 222
# ...of the previous two non-trivial   : 183
# Contextual simplify-reflections      : 5
# Paramodulations                      : 222
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 131
#    Positive orientable unit clauses  : 34
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 90
# Current number of unprocessed clauses: 22
# ...number of literals in the above   : 182
# Current number of archived formulas  : 0
# Current number of archived clauses   : 54
# Clause-clause subsumption calls (NU) : 5292
# Rec. Clause-clause subsumption calls : 266
# Non-unit clause-clause subsumptions  : 66
# Unit Clause-clause subsumption calls : 124
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 12633

# -------------------------------------------------
# User time                : 0.017 s
# System time              : 0.000 s
# Total time               : 0.017 s
# Maximum resident set size: 28712 pages
