# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/6911ccb086fa455bacb7d001c19d5b05.E_input', t4_subset)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/6911ccb086fa455bacb7d001c19d5b05.E_input', t3_subset)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/6911ccb086fa455bacb7d001c19d5b05.E_input', t2_subset)).
fof(t79_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>r1_tarski(k1_zfmisc_1(X1),k1_zfmisc_1(X2))), file('data/MPTP/proofs/6911ccb086fa455bacb7d001c19d5b05.E_input', t79_zfmisc_1)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/6911ccb086fa455bacb7d001c19d5b05.E_input', fc1_subset_1)).
fof(t39_pre_topc, conjecture, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/6911ccb086fa455bacb7d001c19d5b05.E_input', t39_pre_topc)).
fof(d9_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>(m1_pre_topc(X2,X1)<=>(r1_tarski(k2_struct_0(X2),k2_struct_0(X1))&![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>(r2_hidden(X3,u1_pre_topc(X2))<=>?[X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))&(r2_hidden(X4,u1_pre_topc(X1))&X3=k9_subset_1(u1_struct_0(X2),X4,k2_struct_0(X2)))))))))), file('data/MPTP/proofs/6911ccb086fa455bacb7d001c19d5b05.E_input', d9_pre_topc)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/6911ccb086fa455bacb7d001c19d5b05.E_input', dt_m1_pre_topc)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/6911ccb086fa455bacb7d001c19d5b05.E_input', d3_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/6911ccb086fa455bacb7d001c19d5b05.E_input', dt_l1_pre_topc)).
fof(c_0_10, plain, ![X22, X23, X24]:(~r2_hidden(X22,X23)|~m1_subset_1(X23,k1_zfmisc_1(X24))|m1_subset_1(X22,X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_11, plain, ![X20, X21]:((~m1_subset_1(X20,k1_zfmisc_1(X21))|r1_tarski(X20,X21))&(~r1_tarski(X20,X21)|m1_subset_1(X20,k1_zfmisc_1(X21)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_12, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_13, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_14, plain, ![X18, X19]:(~m1_subset_1(X18,X19)|(v1_xboole_0(X19)|r2_hidden(X18,X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_15, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_16, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_17, plain, ![X26, X27]:(~r1_tarski(X26,X27)|r1_tarski(k1_zfmisc_1(X26),k1_zfmisc_1(X27))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t79_zfmisc_1])])).
fof(c_0_18, plain, ![X8]:~v1_xboole_0(k1_zfmisc_1(X8)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
fof(c_0_19, negated_conjecture, ~(![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))))), inference(assume_negation,[status(cth)],[t39_pre_topc])).
fof(c_0_20, plain, ![X9, X10, X11, X13, X15]:(((r1_tarski(k2_struct_0(X10),k2_struct_0(X9))|~m1_pre_topc(X10,X9)|~l1_pre_topc(X10)|~l1_pre_topc(X9))&(((m1_subset_1(esk4_3(X9,X10,X11),k1_zfmisc_1(u1_struct_0(X9)))|~r2_hidden(X11,u1_pre_topc(X10))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~m1_pre_topc(X10,X9)|~l1_pre_topc(X10)|~l1_pre_topc(X9))&((r2_hidden(esk4_3(X9,X10,X11),u1_pre_topc(X9))|~r2_hidden(X11,u1_pre_topc(X10))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~m1_pre_topc(X10,X9)|~l1_pre_topc(X10)|~l1_pre_topc(X9))&(X11=k9_subset_1(u1_struct_0(X10),esk4_3(X9,X10,X11),k2_struct_0(X10))|~r2_hidden(X11,u1_pre_topc(X10))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~m1_pre_topc(X10,X9)|~l1_pre_topc(X10)|~l1_pre_topc(X9))))&(~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X9)))|(~r2_hidden(X13,u1_pre_topc(X9))|X11!=k9_subset_1(u1_struct_0(X10),X13,k2_struct_0(X10)))|r2_hidden(X11,u1_pre_topc(X10))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~m1_pre_topc(X10,X9)|~l1_pre_topc(X10)|~l1_pre_topc(X9))))&((m1_subset_1(esk5_2(X9,X10),k1_zfmisc_1(u1_struct_0(X10)))|~r1_tarski(k2_struct_0(X10),k2_struct_0(X9))|m1_pre_topc(X10,X9)|~l1_pre_topc(X10)|~l1_pre_topc(X9))&((~r2_hidden(esk5_2(X9,X10),u1_pre_topc(X10))|(~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X9)))|(~r2_hidden(X15,u1_pre_topc(X9))|esk5_2(X9,X10)!=k9_subset_1(u1_struct_0(X10),X15,k2_struct_0(X10))))|~r1_tarski(k2_struct_0(X10),k2_struct_0(X9))|m1_pre_topc(X10,X9)|~l1_pre_topc(X10)|~l1_pre_topc(X9))&((m1_subset_1(esk6_2(X9,X10),k1_zfmisc_1(u1_struct_0(X9)))|r2_hidden(esk5_2(X9,X10),u1_pre_topc(X10))|~r1_tarski(k2_struct_0(X10),k2_struct_0(X9))|m1_pre_topc(X10,X9)|~l1_pre_topc(X10)|~l1_pre_topc(X9))&((r2_hidden(esk6_2(X9,X10),u1_pre_topc(X9))|r2_hidden(esk5_2(X9,X10),u1_pre_topc(X10))|~r1_tarski(k2_struct_0(X10),k2_struct_0(X9))|m1_pre_topc(X10,X9)|~l1_pre_topc(X10)|~l1_pre_topc(X9))&(esk5_2(X9,X10)=k9_subset_1(u1_struct_0(X10),esk6_2(X9,X10),k2_struct_0(X10))|r2_hidden(esk5_2(X9,X10),u1_pre_topc(X10))|~r1_tarski(k2_struct_0(X10),k2_struct_0(X9))|m1_pre_topc(X10,X9)|~l1_pre_topc(X10)|~l1_pre_topc(X9))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_pre_topc])])])])])).
fof(c_0_21, plain, ![X28, X29]:(~l1_pre_topc(X28)|(~m1_pre_topc(X29,X28)|l1_pre_topc(X29))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_22, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~r1_tarski(X1,X3)|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_23, plain, (r1_tarski(k1_zfmisc_1(X1),k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_25, negated_conjecture, (l1_pre_topc(esk1_0)&(m1_pre_topc(esk2_0,esk1_0)&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])).
cnf(c_0_26, plain, (r1_tarski(k2_struct_0(X1),k2_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X1)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_28, plain, ![X17]:(~l1_struct_0(X17)|k2_struct_0(X17)=u1_struct_0(X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_29, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24])).
cnf(c_0_30, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, plain, (r1_tarski(k2_struct_0(X1),k2_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(X1))|~r1_tarski(u1_struct_0(esk2_0),X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_34, plain, (r1_tarski(u1_struct_0(X1),k2_struct_0(X2))|~l1_struct_0(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_35, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_36, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_struct_0(X1)))|~l1_struct_0(esk2_0)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
fof(c_0_37, plain, ![X25]:(~l1_pre_topc(X25)|l1_struct_0(X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_38, negated_conjecture, (~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_39, negated_conjecture, (r1_tarski(esk3_0,k2_struct_0(X1))|~l1_struct_0(esk2_0)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_40, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_41, negated_conjecture, (~r1_tarski(esk3_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_38, c_0_13])).
cnf(c_0_42, negated_conjecture, (r1_tarski(esk3_0,u1_struct_0(X1))|~l1_struct_0(esk2_0)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_32]), c_0_40])).
cnf(c_0_43, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_44, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_45, negated_conjecture, (~l1_struct_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), c_0_44])])).
cnf(c_0_46, negated_conjecture, (l1_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_43]), c_0_44])])).
cnf(c_0_47, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_40]), c_0_46])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 48
# Proof object clause steps            : 27
# Proof object formula steps           : 21
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 10
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 67
# ...of these trivial                  : 0
# ...subsumed                          : 7
# ...remaining for further processing  : 60
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 115
# ...of the previous two non-trivial   : 106
# Contextual simplify-reflections      : 7
# Paramodulations                      : 114
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 59
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 50
# Current number of unprocessed clauses: 62
# ...number of literals in the above   : 336
# Current number of archived formulas  : 0
# Current number of archived clauses   : 0
# Clause-clause subsumption calls (NU) : 824
# Rec. Clause-clause subsumption calls : 410
# Non-unit clause-clause subsumptions  : 14
# Unit Clause-clause subsumption calls : 89
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3959

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28744 pages
