# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/696a6c31ebce4c508b209b9b695aab7f.E_input', d1_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/696a6c31ebce4c508b209b9b695aab7f.E_input', d3_tarski)).
fof(t114_zfmisc_1, conjecture, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X2,X1)=>(X1=k1_xboole_0|(X2=k1_xboole_0|X1=X2))), file('data/MPTP/proofs/696a6c31ebce4c508b209b9b695aab7f.E_input', t114_zfmisc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/696a6c31ebce4c508b209b9b695aab7f.E_input', d10_xboole_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/696a6c31ebce4c508b209b9b695aab7f.E_input', fc1_xboole_0)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/696a6c31ebce4c508b209b9b695aab7f.E_input', l54_zfmisc_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/696a6c31ebce4c508b209b9b695aab7f.E_input', t113_zfmisc_1)).
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/696a6c31ebce4c508b209b9b695aab7f.E_input', d2_zfmisc_1)).
fof(c_0_8, plain, ![X21, X22, X23]:((~v1_xboole_0(X21)|~r2_hidden(X22,X21))&(r2_hidden(esk4_1(X23),X23)|v1_xboole_0(X23))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_9, plain, ![X15, X16, X17, X18, X19]:((~r1_tarski(X15,X16)|(~r2_hidden(X17,X15)|r2_hidden(X17,X16)))&((r2_hidden(esk3_2(X18,X19),X18)|r1_tarski(X18,X19))&(~r2_hidden(esk3_2(X18,X19),X19)|r1_tarski(X18,X19)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_10, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_12, negated_conjecture, ~(![X1, X2]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X2,X1)=>(X1=k1_xboole_0|(X2=k1_xboole_0|X1=X2)))), inference(assume_negation,[status(cth)],[t114_zfmisc_1])).
fof(c_0_13, plain, ![X13, X14]:(((r1_tarski(X13,X14)|X13!=X14)&(r1_tarski(X14,X13)|X13!=X14))&(~r1_tarski(X13,X14)|~r1_tarski(X14,X13)|X13=X14)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_14, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
fof(c_0_16, plain, (~epred2_0<=>![X4, X3]:~r2_hidden(X3,X4)), introduced(definition)).
fof(c_0_17, plain, ![X9, X10, X11, X12]:(((r2_hidden(X9,X11)|~r2_hidden(k4_tarski(X9,X10),k2_zfmisc_1(X11,X12)))&(r2_hidden(X10,X12)|~r2_hidden(k4_tarski(X9,X10),k2_zfmisc_1(X11,X12))))&(~r2_hidden(X9,X11)|~r2_hidden(X10,X12)|r2_hidden(k4_tarski(X9,X10),k2_zfmisc_1(X11,X12)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
fof(c_0_18, plain, ![X25, X26]:((k2_zfmisc_1(X25,X26)!=k1_xboole_0|(X25=k1_xboole_0|X26=k1_xboole_0))&((X25!=k1_xboole_0|k2_zfmisc_1(X25,X26)=k1_xboole_0)&(X26!=k1_xboole_0|k2_zfmisc_1(X25,X26)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_19, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k2_zfmisc_1(esk2_0,esk1_0)&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&esk1_0!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
cnf(c_0_20, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_22, plain, (epred2_0|~r2_hidden(X1,X2)), inference(split_equiv,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (X1=k1_xboole_0|X2=k1_xboole_0|k2_zfmisc_1(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k2_zfmisc_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_29, plain, (epred2_0|r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_22, c_0_11])).
fof(c_0_30, plain, (~epred1_0<=>![X1, X2]:(~r2_hidden(X1,X2)|X2!=k1_xboole_0)), introduced(definition)).
cnf(c_0_31, plain, (~r2_hidden(X1,X2)|~r2_hidden(X3,X4)|~v1_xboole_0(k2_zfmisc_1(X4,X2))), inference(spm,[status(thm)],[c_0_10, c_0_23])).
cnf(c_0_32, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_33, negated_conjecture, (k2_zfmisc_1(esk2_0,esk1_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26]), c_0_27])).
cnf(c_0_34, plain, (X1=k1_xboole_0|epred2_0), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35, plain, (~epred2_0|~epred1_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_15])]), c_0_30]), c_0_16])).
cnf(c_0_36, negated_conjecture, (epred2_0), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_37, plain, (~epred1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_36])])).
fof(c_0_38, plain, ![X27, X28, X29, X30, X33, X34, X35, X36, X37, X38, X40, X41]:((((r2_hidden(esk5_4(X27,X28,X29,X30),X27)|~r2_hidden(X30,X29)|X29!=k2_zfmisc_1(X27,X28))&((r2_hidden(esk6_4(X27,X28,X29,X30),X28)|~r2_hidden(X30,X29)|X29!=k2_zfmisc_1(X27,X28))&(X30=k4_tarski(esk5_4(X27,X28,X29,X30),esk6_4(X27,X28,X29,X30))|~r2_hidden(X30,X29)|X29!=k2_zfmisc_1(X27,X28))))&(~r2_hidden(X34,X27)|(~r2_hidden(X35,X28)|X33!=k4_tarski(X34,X35))|r2_hidden(X33,X29)|X29!=k2_zfmisc_1(X27,X28)))&((~r2_hidden(esk7_3(X36,X37,X38),X38)|(~r2_hidden(X40,X36)|(~r2_hidden(X41,X37)|esk7_3(X36,X37,X38)!=k4_tarski(X40,X41)))|X38=k2_zfmisc_1(X36,X37))&((r2_hidden(esk8_3(X36,X37,X38),X36)|r2_hidden(esk7_3(X36,X37,X38),X38)|X38=k2_zfmisc_1(X36,X37))&((r2_hidden(esk9_3(X36,X37,X38),X37)|r2_hidden(esk7_3(X36,X37,X38),X38)|X38=k2_zfmisc_1(X36,X37))&(esk7_3(X36,X37,X38)=k4_tarski(esk8_3(X36,X37,X38),esk9_3(X36,X37,X38))|r2_hidden(esk7_3(X36,X37,X38),X38)|X38=k2_zfmisc_1(X36,X37)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
cnf(c_0_39, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,X2))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_40, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(sr,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_30]), c_0_37])).
cnf(c_0_41, plain, (r2_hidden(esk8_3(X1,X2,X3),X1)|r2_hidden(esk7_3(X1,X2,X3),X3)|X3=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_42, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(k4_tarski(X2,X1),k2_zfmisc_1(esk2_0,esk1_0))), inference(spm,[status(thm)],[c_0_39, c_0_25])).
cnf(c_0_43, plain, (X1=k2_zfmisc_1(X2,X3)|r2_hidden(esk8_3(X2,X3,X1),X2)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_44, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)|~r2_hidden(X2,esk2_0)), inference(spm,[status(thm)],[c_0_42, c_0_23])).
cnf(c_0_45, plain, (r2_hidden(esk4_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_46, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,k1_xboole_0),X1)), inference(er,[status(thm)],[c_0_43])).
cnf(c_0_47, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_48, negated_conjecture, (r2_hidden(esk4_1(esk1_0),esk2_0)|v1_xboole_0(esk1_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_49, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_10, c_0_46])).
cnf(c_0_50, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(esk2_0,esk1_0))), inference(spm,[status(thm)],[c_0_47, c_0_25])).
cnf(c_0_51, negated_conjecture, (r2_hidden(esk4_1(esk1_0),esk2_0)|v1_xboole_0(esk2_0)|v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_48, c_0_45])).
cnf(c_0_52, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_33, c_0_49])).
cnf(c_0_53, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X2,esk1_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_50, c_0_23])).
cnf(c_0_54, negated_conjecture, (r2_hidden(esk4_1(esk1_0),esk2_0)|v1_xboole_0(esk1_0)), inference(sr,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_55, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_49]), c_0_33])).
cnf(c_0_56, negated_conjecture, (r2_hidden(X1,esk1_0)|v1_xboole_0(esk1_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_53, c_0_45])).
cnf(c_0_57, negated_conjecture, (r2_hidden(esk4_1(esk1_0),esk2_0)), inference(sr,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_58, plain, (r2_hidden(esk9_3(X1,X2,X3),X2)|r2_hidden(esk7_3(X1,X2,X3),X3)|X3=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_59, negated_conjecture, (r2_hidden(esk4_1(esk1_0),esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_55])).
cnf(c_0_60, negated_conjecture, (X1=k2_zfmisc_1(esk2_0,esk1_0)|r2_hidden(esk9_3(esk1_0,esk2_0,X1),esk2_0)|r2_hidden(esk7_3(esk1_0,esk2_0,X1),X1)), inference(spm,[status(thm)],[c_0_25, c_0_58])).
cnf(c_0_61, plain, (r1_tarski(X1,X2)|~r2_hidden(esk3_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_62, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_53, c_0_59])).
cnf(c_0_63, negated_conjecture, (X1=k2_zfmisc_1(esk2_0,esk1_0)|r2_hidden(esk9_3(esk1_0,esk2_0,X1),esk2_0)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_40, c_0_60])).
cnf(c_0_64, negated_conjecture, (r1_tarski(X1,esk1_0)|~r2_hidden(esk3_2(X1,esk1_0),esk2_0)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_65, negated_conjecture, (r1_tarski(esk1_0,X1)|r2_hidden(esk3_2(esk1_0,X1),esk2_0)|~r2_hidden(X2,esk2_0)), inference(spm,[status(thm)],[c_0_44, c_0_11])).
cnf(c_0_66, negated_conjecture, (r2_hidden(esk9_3(esk1_0,esk2_0,k1_xboole_0),esk2_0)), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_63]), c_0_33])).
cnf(c_0_67, negated_conjecture, (r1_tarski(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_64, c_0_11])).
cnf(c_0_68, negated_conjecture, (esk1_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_69, negated_conjecture, (r1_tarski(esk1_0,X1)|r2_hidden(esk3_2(esk1_0,X1),esk2_0)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_70, negated_conjecture, (~r1_tarski(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_67]), c_0_68])).
cnf(c_0_71, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_69]), c_0_70]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 72
# Proof object clause steps            : 54
# Proof object formula steps           : 18
# Proof object conjectures             : 31
# Proof object clause conjectures      : 28
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 32
# Proof object simplifying inferences  : 16
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 1169
# ...of these trivial                  : 6
# ...subsumed                          : 877
# ...remaining for further processing  : 286
# Other redundant clauses eliminated   : 267
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 19
# Backward-rewritten                   : 13
# Generated clauses                    : 13989
# ...of the previous two non-trivial   : 12799
# Contextual simplify-reflections      : 11
# Paramodulations                      : 13665
# Factorizations                       : 15
# Equation resolutions                 : 301
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 222
#    Positive orientable unit clauses  : 18
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 10
#    Non-unit-clauses                  : 194
# Current number of unprocessed clauses: 11509
# ...number of literals in the above   : 71338
# Current number of archived formulas  : 0
# Current number of archived clauses   : 60
# Clause-clause subsumption calls (NU) : 13461
# Rec. Clause-clause subsumption calls : 7124
# Non-unit clause-clause subsumptions  : 560
# Unit Clause-clause subsumption calls : 209
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12
# BW rewrite match successes           : 8
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 230634

# -------------------------------------------------
# User time                : 0.185 s
# System time              : 0.010 s
# Total time               : 0.195 s
# Maximum resident set size: 28600 pages
