# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', t100_xboole_1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', t12_setfam_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', t3_boole)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', t2_boole)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', t48_xboole_1)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', d3_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', dt_l1_pre_topc)).
fof(t54_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k2_pre_topc(X1,X2))<=>(~(v2_struct_0(X1))&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>~((v3_pre_topc(X4,X1)&(r2_hidden(X3,X4)&r1_xboole_0(X2,X4)))))))))), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', t54_pre_topc)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', t4_subset)).
fof(t27_tops_1, axiom, ![X1]:(l1_pre_topc(X1)=>k2_pre_topc(X1,k2_struct_0(X1))=k2_struct_0(X1)), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', t27_tops_1)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', d4_subset_1)).
fof(d5_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', d5_subset_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', t28_xboole_1)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', l3_subset_1)).
fof(t29_tops_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v4_pre_topc(X2,X1)<=>v3_pre_topc(k3_subset_1(u1_struct_0(X1),X2),X1)))), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', t29_tops_1)).
fof(t80_tops_1, conjecture, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v1_tops_1(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>~((~(X3=k1_xboole_0)&(v3_pre_topc(X3,X1)&r1_xboole_0(X2,X3)))))))), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', t80_tops_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', t3_subset)).
fof(t52_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v4_pre_topc(X2,X1)=>k2_pre_topc(X1,X2)=X2)&((v2_pre_topc(X1)&k2_pre_topc(X1,X2)=X2)=>v4_pre_topc(X2,X1))))), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', t52_pre_topc)).
fof(d13_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k2_pre_topc(X1,X2)<=>![X4]:(r2_hidden(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))=>~((v3_pre_topc(X5,X1)&(r2_hidden(X4,X5)&r1_xboole_0(X2,X5))))))))))), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', d13_pre_topc)).
fof(d3_tops_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v1_tops_1(X2,X1)<=>k2_pre_topc(X1,X2)=k2_struct_0(X1)))), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', d3_tops_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', d7_xboole_0)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', t4_subset_1)).
fof(fc10_tops_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>v3_pre_topc(k2_struct_0(X1),X1)), file('data/MPTP/proofs/698b721337044deea8fd6de2178995f1.E_input', fc10_tops_1)).
fof(c_0_24, plain, ![X55, X56]:k4_xboole_0(X55,X56)=k5_xboole_0(X55,k3_xboole_0(X55,X56)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_25, plain, ![X15, X16]:k1_setfam_1(k2_tarski(X15,X16))=k3_xboole_0(X15,X16), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_26, plain, ![X49]:k4_xboole_0(X49,k1_xboole_0)=X49, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_27, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_28, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_29, plain, ![X12]:k3_xboole_0(X12,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_30, plain, ![X47, X48]:k4_xboole_0(X47,k4_xboole_0(X47,X48))=k3_xboole_0(X47,X48), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_31, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X2)))), inference(rw,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_33, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_35, plain, (k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,k1_xboole_0)))=X1), inference(rw,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_36, plain, (k1_setfam_1(k2_tarski(X1,k1_xboole_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_33, c_0_28])).
fof(c_0_37, plain, ![X13]:(~l1_struct_0(X13)|k2_struct_0(X13)=u1_struct_0(X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
fof(c_0_38, plain, ![X58]:(~l1_pre_topc(X58)|l1_struct_0(X58)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_39, plain, (k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X2))))))=k1_setfam_1(k2_tarski(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_28]), c_0_32]), c_0_32])).
cnf(c_0_40, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_35, c_0_36])).
fof(c_0_41, plain, ![X39, X40, X41, X42]:(((~v2_struct_0(X39)|~r2_hidden(X41,k2_pre_topc(X39,X40))|~m1_subset_1(X41,u1_struct_0(X39))|~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X39)))|~l1_pre_topc(X39))&(~m1_subset_1(X42,k1_zfmisc_1(u1_struct_0(X39)))|(~v3_pre_topc(X42,X39)|(~r2_hidden(X41,X42)|~r1_xboole_0(X40,X42)))|~r2_hidden(X41,k2_pre_topc(X39,X40))|~m1_subset_1(X41,u1_struct_0(X39))|~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X39)))|~l1_pre_topc(X39)))&((m1_subset_1(esk7_3(X39,X40,X41),k1_zfmisc_1(u1_struct_0(X39)))|v2_struct_0(X39)|r2_hidden(X41,k2_pre_topc(X39,X40))|~m1_subset_1(X41,u1_struct_0(X39))|~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X39)))|~l1_pre_topc(X39))&((v3_pre_topc(esk7_3(X39,X40,X41),X39)|v2_struct_0(X39)|r2_hidden(X41,k2_pre_topc(X39,X40))|~m1_subset_1(X41,u1_struct_0(X39))|~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X39)))|~l1_pre_topc(X39))&((r2_hidden(X41,esk7_3(X39,X40,X41))|v2_struct_0(X39)|r2_hidden(X41,k2_pre_topc(X39,X40))|~m1_subset_1(X41,u1_struct_0(X39))|~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X39)))|~l1_pre_topc(X39))&(r1_xboole_0(X40,esk7_3(X39,X40,X41))|v2_struct_0(X39)|r2_hidden(X41,k2_pre_topc(X39,X40))|~m1_subset_1(X41,u1_struct_0(X39))|~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X39)))|~l1_pre_topc(X39)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t54_pre_topc])])])])])])).
fof(c_0_42, plain, ![X52, X53, X54]:(~r2_hidden(X52,X53)|~m1_subset_1(X53,k1_zfmisc_1(X54))|m1_subset_1(X52,X54)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_43, plain, ![X46]:(~l1_pre_topc(X46)|k2_pre_topc(X46,k2_struct_0(X46))=k2_struct_0(X46)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t27_tops_1])])).
cnf(c_0_44, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_45, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_46, plain, ![X38]:m1_subset_1(k2_subset_1(X38),k1_zfmisc_1(X38)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_47, plain, ![X14]:k2_subset_1(X14)=X14, inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_48, plain, ![X29, X30]:(~m1_subset_1(X30,k1_zfmisc_1(X29))|k3_subset_1(X29,X30)=k4_xboole_0(X29,X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
fof(c_0_49, plain, ![X33, X34]:(~r1_tarski(X33,X34)|k3_xboole_0(X33,X34)=X33), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_50, plain, (k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X1)))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_36]), c_0_40])).
cnf(c_0_51, plain, (~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v3_pre_topc(X1,X2)|~r2_hidden(X3,X1)|~r1_xboole_0(X4,X1)|~r2_hidden(X3,k2_pre_topc(X2,X4))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_52, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_53, plain, (k2_pre_topc(X1,k2_struct_0(X1))=k2_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_54, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_55, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_56, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_57, plain, ![X17, X18, X19]:(~m1_subset_1(X18,k1_zfmisc_1(X17))|(~r2_hidden(X19,X18)|r2_hidden(X19,X17))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
cnf(c_0_58, plain, (k3_subset_1(X2,X1)=k4_xboole_0(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_60, plain, (k1_setfam_1(k2_tarski(X1,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_50]), c_0_36]), c_0_40])).
cnf(c_0_61, plain, (~r2_hidden(X1,k2_pre_topc(X2,X3))|~r2_hidden(X1,X4)|~r1_xboole_0(X3,X4)|~v3_pre_topc(X4,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(csr,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_62, plain, (k2_pre_topc(X1,u1_struct_0(X1))=u1_struct_0(X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_63, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_64, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
fof(c_0_65, plain, ![X31, X32]:((~v4_pre_topc(X32,X31)|v3_pre_topc(k3_subset_1(u1_struct_0(X31),X32),X31)|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|~l1_pre_topc(X31))&(~v3_pre_topc(k3_subset_1(u1_struct_0(X31),X32),X31)|v4_pre_topc(X32,X31)|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|~l1_pre_topc(X31))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t29_tops_1])])])])).
cnf(c_0_66, plain, (k5_xboole_0(X2,k1_setfam_1(k2_tarski(X2,X1)))=k3_subset_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_58, c_0_32])).
cnf(c_0_67, plain, (k1_setfam_1(k2_tarski(X1,X2))=X1|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_59, c_0_28])).
cnf(c_0_68, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[c_0_50, c_0_60])).
fof(c_0_69, negated_conjecture, ~(![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v1_tops_1(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>~((~(X3=k1_xboole_0)&(v3_pre_topc(X3,X1)&r1_xboole_0(X2,X3))))))))), inference(assume_negation,[status(cth)],[t80_tops_1])).
cnf(c_0_70, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(u1_struct_0(X3),X2)|~v3_pre_topc(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63])]), c_0_64])).
cnf(c_0_71, plain, (r1_xboole_0(X1,esk7_3(X2,X1,X3))|v2_struct_0(X2)|r2_hidden(X3,k2_pre_topc(X2,X1))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_72, plain, (v3_pre_topc(k3_subset_1(u1_struct_0(X2),X1),X2)|~v4_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_73, plain, (k3_subset_1(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68])).
fof(c_0_74, plain, ![X10, X11]:((~m1_subset_1(X10,k1_zfmisc_1(X11))|r1_tarski(X10,X11))&(~r1_tarski(X10,X11)|m1_subset_1(X10,k1_zfmisc_1(X11)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_75, plain, ![X36, X37]:((~v4_pre_topc(X37,X36)|k2_pre_topc(X36,X37)=X37|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X36)))|~l1_pre_topc(X36))&(~v2_pre_topc(X36)|k2_pre_topc(X36,X37)!=X37|v4_pre_topc(X37,X36)|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X36)))|~l1_pre_topc(X36))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t52_pre_topc])])])])).
fof(c_0_76, negated_conjecture, ![X9]:((v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(((m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_tops_1(esk2_0,esk1_0))&((esk3_0!=k1_xboole_0|~v1_tops_1(esk2_0,esk1_0))&((v3_pre_topc(esk3_0,esk1_0)|~v1_tops_1(esk2_0,esk1_0))&(r1_xboole_0(esk2_0,esk3_0)|~v1_tops_1(esk2_0,esk1_0)))))&(v1_tops_1(esk2_0,esk1_0)|(~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(esk1_0)))|(X9=k1_xboole_0|(~v3_pre_topc(X9,esk1_0)|~r1_xboole_0(esk2_0,X9)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_69])])])])])])).
cnf(c_0_77, plain, (v2_struct_0(X1)|r2_hidden(X2,k2_pre_topc(X1,u1_struct_0(X3)))|~r2_hidden(X4,esk7_3(X1,u1_struct_0(X3),X2))|~v3_pre_topc(esk7_3(X1,u1_struct_0(X3),X2),X3)|~m1_subset_1(esk7_3(X1,u1_struct_0(X3),X2),k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(u1_struct_0(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X3)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_78, plain, (v3_pre_topc(esk7_3(X1,X2,X3),X1)|v2_struct_0(X1)|r2_hidden(X3,k2_pre_topc(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_79, plain, (~v2_struct_0(X1)|~r2_hidden(X2,k2_pre_topc(X1,X3))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_80, plain, (v4_pre_topc(X2,X1)|~v3_pre_topc(k3_subset_1(u1_struct_0(X1),X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_81, plain, (v3_pre_topc(k1_xboole_0,X1)|~v4_pre_topc(X2,X1)|~r1_tarski(u1_struct_0(X1),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_82, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_83, plain, (v4_pre_topc(X2,X1)|~v2_pre_topc(X1)|k2_pre_topc(X1,X2)!=X2|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_84, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_85, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_86, plain, (v2_struct_0(X1)|r2_hidden(X2,k2_pre_topc(X1,u1_struct_0(X1)))|~r2_hidden(X3,esk7_3(X1,u1_struct_0(X1),X2))|~m1_subset_1(esk7_3(X1,u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_63])])).
cnf(c_0_87, plain, (m1_subset_1(esk7_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|r2_hidden(X3,k2_pre_topc(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_88, plain, ![X20, X21, X22, X23, X24, X28]:(((~r2_hidden(X23,X22)|(~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X20)))|(~v3_pre_topc(X24,X20)|(~r2_hidden(X23,X24)|~r1_xboole_0(X21,X24))))|~r2_hidden(X23,u1_struct_0(X20))|X22!=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))&((m1_subset_1(esk4_4(X20,X21,X22,X23),k1_zfmisc_1(u1_struct_0(X20)))|r2_hidden(X23,X22)|~r2_hidden(X23,u1_struct_0(X20))|X22!=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))&((v3_pre_topc(esk4_4(X20,X21,X22,X23),X20)|r2_hidden(X23,X22)|~r2_hidden(X23,u1_struct_0(X20))|X22!=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))&((r2_hidden(X23,esk4_4(X20,X21,X22,X23))|r2_hidden(X23,X22)|~r2_hidden(X23,u1_struct_0(X20))|X22!=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))&(r1_xboole_0(X21,esk4_4(X20,X21,X22,X23))|r2_hidden(X23,X22)|~r2_hidden(X23,u1_struct_0(X20))|X22!=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))))))&((r2_hidden(esk5_3(X20,X21,X22),u1_struct_0(X20))|X22=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))&(((m1_subset_1(esk6_3(X20,X21,X22),k1_zfmisc_1(u1_struct_0(X20)))|~r2_hidden(esk5_3(X20,X21,X22),X22)|X22=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))&((v3_pre_topc(esk6_3(X20,X21,X22),X20)|~r2_hidden(esk5_3(X20,X21,X22),X22)|X22=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))&((r2_hidden(esk5_3(X20,X21,X22),esk6_3(X20,X21,X22))|~r2_hidden(esk5_3(X20,X21,X22),X22)|X22=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))&(r1_xboole_0(X21,esk6_3(X20,X21,X22))|~r2_hidden(esk5_3(X20,X21,X22),X22)|X22=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20)))))&(r2_hidden(esk5_3(X20,X21,X22),X22)|(~m1_subset_1(X28,k1_zfmisc_1(u1_struct_0(X20)))|(~v3_pre_topc(X28,X20)|(~r2_hidden(esk5_3(X20,X21,X22),X28)|~r1_xboole_0(X21,X28))))|X22=k2_pre_topc(X20,X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~l1_pre_topc(X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d13_pre_topc])])])])])).
cnf(c_0_89, plain, (~v2_struct_0(X1)|~r2_hidden(X2,k2_struct_0(X1))|~m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_53]), c_0_52])).
cnf(c_0_90, plain, (v4_pre_topc(X1,X2)|~r1_tarski(u1_struct_0(X2),X1)|~v3_pre_topc(k1_xboole_0,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_80, c_0_73])).
cnf(c_0_91, plain, (v3_pre_topc(k1_xboole_0,X1)|~v4_pre_topc(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(X2))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_92, negated_conjecture, (v4_pre_topc(X1,esk1_0)|k2_pre_topc(esk1_0,X1)!=X1|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_85])])).
cnf(c_0_93, plain, (v2_struct_0(X1)|r2_hidden(X2,k2_pre_topc(X1,u1_struct_0(X1)))|~r2_hidden(X3,esk7_3(X1,u1_struct_0(X1),X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_63])])).
cnf(c_0_94, plain, (r2_hidden(X1,esk7_3(X2,X3,X1))|v2_struct_0(X2)|r2_hidden(X1,k2_pre_topc(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_95, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_52, c_0_63])).
cnf(c_0_96, plain, (r2_hidden(esk5_3(X1,X2,X3),u1_struct_0(X1))|X3=k2_pre_topc(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_97, plain, (~v2_struct_0(X1)|~r2_hidden(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_54]), c_0_63])])).
cnf(c_0_98, plain, (k2_pre_topc(X2,X1)=X1|~v4_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_99, plain, (v4_pre_topc(X1,X2)|~v3_pre_topc(k1_xboole_0,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(X1))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_90, c_0_82])).
cnf(c_0_100, negated_conjecture, (v3_pre_topc(k1_xboole_0,esk1_0)|k2_pre_topc(esk1_0,X1)!=X1|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_85])])).
fof(c_0_101, plain, ![X44, X45]:((~v1_tops_1(X45,X44)|k2_pre_topc(X44,X45)=k2_struct_0(X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X44)))|~l1_pre_topc(X44))&(k2_pre_topc(X44,X45)!=k2_struct_0(X44)|v1_tops_1(X45,X44)|~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X44)))|~l1_pre_topc(X44))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tops_1])])])])).
cnf(c_0_102, plain, (v2_struct_0(X1)|r2_hidden(X2,k2_pre_topc(X1,u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_63])])).
cnf(c_0_103, plain, (X1=k2_pre_topc(X2,X3)|m1_subset_1(esk5_3(X2,X3,X1),u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_104, plain, (X1=k2_pre_topc(X2,X3)|~v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_97, c_0_96])).
cnf(c_0_105, plain, (k2_pre_topc(X1,X2)=X2|~v3_pre_topc(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(X2))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_106, negated_conjecture, (v3_pre_topc(k1_xboole_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_62]), c_0_63]), c_0_85])])).
cnf(c_0_107, plain, (v1_tops_1(X2,X1)|k2_pre_topc(X1,X2)!=k2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_108, plain, (X1=k2_pre_topc(X2,X3)|r2_hidden(esk5_3(X2,X3,X1),k2_pre_topc(X2,u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_104])).
cnf(c_0_109, negated_conjecture, (k2_pre_topc(esk1_0,X1)=X1|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_85])])).
cnf(c_0_110, negated_conjecture, (r1_xboole_0(esk2_0,esk3_0)|~v1_tops_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_111, plain, (v1_tops_1(X1,X2)|k2_pre_topc(X2,X1)!=u1_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_107, c_0_54])).
cnf(c_0_112, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_113, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_tops_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_114, negated_conjecture, (v3_pre_topc(esk3_0,esk1_0)|~v1_tops_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_115, negated_conjecture, (esk3_0!=k1_xboole_0|~v1_tops_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_116, plain, (r1_xboole_0(X1,esk6_3(X2,X1,X3))|X3=k2_pre_topc(X2,X1)|~r2_hidden(esk5_3(X2,X1,X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_117, negated_conjecture, (X1=k2_pre_topc(esk1_0,X2)|r2_hidden(esk5_3(esk1_0,X2,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_85]), c_0_63])])).
cnf(c_0_118, negated_conjecture, (v1_tops_1(esk2_0,esk1_0)|X1=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v3_pre_topc(X1,esk1_0)|~r1_xboole_0(esk2_0,X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_119, negated_conjecture, (r1_xboole_0(esk2_0,esk3_0)|k2_pre_topc(esk1_0,esk2_0)!=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_112]), c_0_85])])).
cnf(c_0_120, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|k2_pre_topc(esk1_0,esk2_0)!=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_111]), c_0_112]), c_0_85])])).
cnf(c_0_121, negated_conjecture, (v3_pre_topc(esk3_0,esk1_0)|k2_pre_topc(esk1_0,esk2_0)!=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_111]), c_0_112]), c_0_85])])).
cnf(c_0_122, negated_conjecture, (k2_pre_topc(esk1_0,esk2_0)!=u1_struct_0(esk1_0)|esk3_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_111]), c_0_112]), c_0_85])])).
cnf(c_0_123, plain, (k2_pre_topc(X2,X1)=k2_struct_0(X2)|~v1_tops_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_124, plain, (v1_tops_1(u1_struct_0(X1),X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_62]), c_0_63])]), c_0_54])).
fof(c_0_125, plain, ![X50, X51]:((~r1_xboole_0(X50,X51)|k3_xboole_0(X50,X51)=k1_xboole_0)&(k3_xboole_0(X50,X51)!=k1_xboole_0|r1_xboole_0(X50,X51))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_126, negated_conjecture, (k2_pre_topc(esk1_0,X1)=u1_struct_0(esk1_0)|r1_xboole_0(X1,esk6_3(esk1_0,X1,u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_63]), c_0_85])])).
cnf(c_0_127, negated_conjecture, (v1_tops_1(esk2_0,esk1_0)|k2_pre_topc(esk1_0,esk2_0)!=u1_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_120]), c_0_121]), c_0_122])).
cnf(c_0_128, plain, (v3_pre_topc(esk6_3(X1,X2,X3),X1)|X3=k2_pre_topc(X1,X2)|~r2_hidden(esk5_3(X1,X2,X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_129, plain, (~r2_hidden(X1,k2_struct_0(X2))|~r2_hidden(X1,X3)|~r1_xboole_0(k2_struct_0(X2),X3)|~v3_pre_topc(X3,X2)|~m1_subset_1(k2_struct_0(X2),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_61, c_0_53])).
cnf(c_0_130, plain, (k2_pre_topc(X1,u1_struct_0(X1))=k2_struct_0(X1)|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124]), c_0_63])])).
cnf(c_0_131, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_125])).
cnf(c_0_132, negated_conjecture, (esk6_3(esk1_0,esk2_0,u1_struct_0(esk1_0))=k1_xboole_0|v1_tops_1(esk2_0,esk1_0)|~v3_pre_topc(esk6_3(esk1_0,esk2_0,u1_struct_0(esk1_0)),esk1_0)|~m1_subset_1(esk6_3(esk1_0,esk2_0,u1_struct_0(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_126]), c_0_112])]), c_0_127])).
cnf(c_0_133, plain, (k2_pre_topc(X1,X2)=u1_struct_0(X1)|v3_pre_topc(esk6_3(X1,X2,u1_struct_0(X1)),X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_96]), c_0_63])])).
cnf(c_0_134, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|X3=k2_pre_topc(X1,X2)|~r2_hidden(esk5_3(X1,X2,X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_135, plain, (~r2_hidden(X1,k2_pre_topc(X2,u1_struct_0(X2)))|~r2_hidden(X1,X3)|~r1_xboole_0(k2_pre_topc(X2,u1_struct_0(X2)),X3)|~v3_pre_topc(X3,X2)|~m1_subset_1(k2_pre_topc(X2,u1_struct_0(X2)),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_129, c_0_130])).
cnf(c_0_136, plain, (r1_xboole_0(X1,X2)|k1_setfam_1(k2_tarski(X1,X2))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_131, c_0_28])).
fof(c_0_137, plain, ![X35]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X35)), inference(variable_rename,[status(thm)],[t4_subset_1])).
cnf(c_0_138, negated_conjecture, (esk6_3(esk1_0,esk2_0,u1_struct_0(esk1_0))=k1_xboole_0|v1_tops_1(esk2_0,esk1_0)|~m1_subset_1(esk6_3(esk1_0,esk2_0,u1_struct_0(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_133]), c_0_112]), c_0_85])]), c_0_127])).
cnf(c_0_139, plain, (k2_pre_topc(X1,X2)=u1_struct_0(X1)|m1_subset_1(esk6_3(X1,X2,u1_struct_0(X1)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_96]), c_0_63])])).
cnf(c_0_140, negated_conjecture, (~r2_hidden(X1,X2)|~r1_xboole_0(u1_struct_0(esk1_0),X2)|~v3_pre_topc(X2,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_109]), c_0_63]), c_0_85])]), c_0_64])).
cnf(c_0_141, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_136, c_0_36])).
cnf(c_0_142, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_137])).
cnf(c_0_143, plain, (r2_hidden(esk5_3(X1,X2,X3),esk6_3(X1,X2,X3))|X3=k2_pre_topc(X1,X2)|~r2_hidden(esk5_3(X1,X2,X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_144, negated_conjecture, (esk6_3(esk1_0,esk2_0,u1_struct_0(esk1_0))=k1_xboole_0|v1_tops_1(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_112]), c_0_85])]), c_0_127])).
cnf(c_0_145, negated_conjecture, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_141]), c_0_106]), c_0_142])])).
cnf(c_0_146, negated_conjecture, (v1_tops_1(k2_struct_0(esk1_0),esk1_0)|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k2_struct_0(esk1_0)))), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_109]), c_0_85])])])).
cnf(c_0_147, plain, (k2_pre_topc(X1,k2_pre_topc(X1,u1_struct_0(X1)))=k2_pre_topc(X1,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_53, c_0_130])).
cnf(c_0_148, negated_conjecture, (v1_tops_1(esk2_0,esk1_0)|~r2_hidden(esk5_3(esk1_0,esk2_0,u1_struct_0(esk1_0)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144]), c_0_63]), c_0_112]), c_0_85])]), c_0_145]), c_0_127])).
cnf(c_0_149, negated_conjecture, (v1_tops_1(u1_struct_0(esk1_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_54]), c_0_63]), c_0_85])])).
cnf(c_0_150, negated_conjecture, (k2_pre_topc(esk1_0,u1_struct_0(esk1_0))=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_109]), c_0_85]), c_0_63])])).
fof(c_0_151, plain, ![X57]:(~v2_pre_topc(X57)|~l1_pre_topc(X57)|v3_pre_topc(k2_struct_0(X57),X57)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc10_tops_1])])).
cnf(c_0_152, negated_conjecture, (v1_tops_1(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_96]), c_0_63]), c_0_112]), c_0_85])]), c_0_127])).
cnf(c_0_153, negated_conjecture, (k2_struct_0(esk1_0)=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_149]), c_0_150]), c_0_63]), c_0_85])])).
cnf(c_0_154, plain, (v3_pre_topc(k2_struct_0(X1),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_151])).
cnf(c_0_155, negated_conjecture, (k2_pre_topc(esk1_0,esk2_0)=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_152]), c_0_153]), c_0_112]), c_0_85])])).
cnf(c_0_156, plain, (v3_pre_topc(u1_struct_0(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_154, c_0_54])).
cnf(c_0_157, plain, (k3_subset_1(X1,k1_xboole_0)=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_36]), c_0_40]), c_0_142])])).
cnf(c_0_158, negated_conjecture, (~r2_hidden(X1,X2)|~r1_xboole_0(esk2_0,X2)|~v3_pre_topc(X2,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_155]), c_0_112]), c_0_85])]), c_0_64])).
cnf(c_0_159, negated_conjecture, (r1_xboole_0(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_110, c_0_152])])).
cnf(c_0_160, negated_conjecture, (v3_pre_topc(esk3_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_114, c_0_152])])).
cnf(c_0_161, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_113, c_0_152])])).
cnf(c_0_162, plain, (r2_hidden(esk5_3(X1,X2,X3),X3)|X3=k2_pre_topc(X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v3_pre_topc(X4,X1)|~r2_hidden(esk5_3(X1,X2,X3),X4)|~r1_xboole_0(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_163, negated_conjecture, (v3_pre_topc(u1_struct_0(esk1_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_84]), c_0_85])])).
cnf(c_0_164, plain, (v4_pre_topc(k1_xboole_0,X1)|~v3_pre_topc(u1_struct_0(X1),X1)|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_157]), c_0_142])])).
cnf(c_0_165, negated_conjecture, (~r2_hidden(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159]), c_0_160]), c_0_161])])).
cnf(c_0_166, negated_conjecture, (X1=k2_pre_topc(esk1_0,X2)|r2_hidden(esk5_3(esk1_0,X2,X1),X1)|~r1_xboole_0(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_117]), c_0_163]), c_0_63]), c_0_85])])).
cnf(c_0_167, plain, (r1_xboole_0(k1_xboole_0,X1)|~r1_tarski(k1_xboole_0,X1)), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_67])])).
cnf(c_0_168, negated_conjecture, (v4_pre_topc(k1_xboole_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_163]), c_0_85])])).
cnf(c_0_169, negated_conjecture, (k2_pre_topc(esk1_0,X1)=esk3_0|~r1_xboole_0(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_166]), c_0_161])])).
cnf(c_0_170, plain, (r1_xboole_0(k1_xboole_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_82]), c_0_142])])).
cnf(c_0_171, negated_conjecture, (k2_pre_topc(esk1_0,k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_168]), c_0_142]), c_0_85])])).
cnf(c_0_172, negated_conjecture, (esk3_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_115, c_0_152])])).
cnf(c_0_173, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_170]), c_0_171]), c_0_142])]), c_0_172]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 174
# Proof object clause steps            : 125
# Proof object formula steps           : 49
# Proof object conjectures             : 46
# Proof object clause conjectures      : 43
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 44
# Proof object initial formulas used   : 24
# Proof object generating inferences   : 66
# Proof object simplifying inferences  : 155
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 24
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 51
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 48
# Processed clauses                    : 1056
# ...of these trivial                  : 61
# ...subsumed                          : 572
# ...remaining for further processing  : 423
# Other redundant clauses eliminated   : 16
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 66
# Backward-rewritten                   : 70
# Generated clauses                    : 1991
# ...of the previous two non-trivial   : 1592
# Contextual simplify-reflections      : 38
# Paramodulations                      : 1971
# Factorizations                       : 0
# Equation resolutions                 : 20
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 235
#    Positive orientable unit clauses  : 28
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 203
# Current number of unprocessed clauses: 549
# ...number of literals in the above   : 3067
# Current number of archived formulas  : 0
# Current number of archived clauses   : 186
# Clause-clause subsumption calls (NU) : 30013
# Rec. Clause-clause subsumption calls : 4906
# Non-unit clause-clause subsumptions  : 599
# Unit Clause-clause subsumption calls : 706
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 30
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 52538

# -------------------------------------------------
# User time                : 0.052 s
# System time              : 0.000 s
# Total time               : 0.052 s
# Maximum resident set size: 28796 pages
