# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/69aa9a991f274eccb321e2d6f9cd0b5a.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/69aa9a991f274eccb321e2d6f9cd0b5a.E_input', t70_enumset1)).
fof(redefinition_k4_subset_1, axiom, ![X1, X2, X3]:((m1_subset_1(X2,k1_zfmisc_1(X1))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>k4_subset_1(X1,X2,X3)=k2_xboole_0(X2,X3)), file('data/MPTP/proofs/69aa9a991f274eccb321e2d6f9cd0b5a.E_input', redefinition_k4_subset_1)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/69aa9a991f274eccb321e2d6f9cd0b5a.E_input', t98_xboole_1)).
fof(t63_setfam_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((v3_setfam_1(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))=>![X3]:((v3_setfam_1(X3,X1)&m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1))))=>v3_setfam_1(k4_subset_1(k1_zfmisc_1(X1),X2,X3),X1)))), file('data/MPTP/proofs/69aa9a991f274eccb321e2d6f9cd0b5a.E_input', t63_setfam_1)).
fof(d13_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>(v3_setfam_1(X2,X1)<=>~(r2_hidden(X1,X2)))), file('data/MPTP/proofs/69aa9a991f274eccb321e2d6f9cd0b5a.E_input', d13_setfam_1)).
fof(dt_k4_subset_1, axiom, ![X1, X2, X3]:((m1_subset_1(X2,k1_zfmisc_1(X1))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>m1_subset_1(k4_subset_1(X1,X2,X3),k1_zfmisc_1(X1))), file('data/MPTP/proofs/69aa9a991f274eccb321e2d6f9cd0b5a.E_input', dt_k4_subset_1)).
fof(t1_xboole_0, axiom, ![X1, X2, X3]:(r2_hidden(X1,k5_xboole_0(X2,X3))<=>~((r2_hidden(X1,X2)<=>r2_hidden(X1,X3)))), file('data/MPTP/proofs/69aa9a991f274eccb321e2d6f9cd0b5a.E_input', t1_xboole_0)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/69aa9a991f274eccb321e2d6f9cd0b5a.E_input', d5_xboole_0)).
fof(c_0_9, plain, ![X5, X6]:k3_tarski(k2_tarski(X5,X6))=k2_xboole_0(X5,X6), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_10, plain, ![X7, X8]:k1_enumset1(X7,X7,X8)=k2_tarski(X7,X8), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_11, plain, ![X9, X10, X11]:(~m1_subset_1(X10,k1_zfmisc_1(X9))|~m1_subset_1(X11,k1_zfmisc_1(X9))|k4_subset_1(X9,X10,X11)=k2_xboole_0(X10,X11)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_subset_1])])).
cnf(c_0_12, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_14, plain, ![X12, X13]:k2_xboole_0(X12,X13)=k5_xboole_0(X12,k4_xboole_0(X13,X12)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
cnf(c_0_15, plain, (k4_subset_1(X2,X1,X3)=k2_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_17, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:((v3_setfam_1(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))=>![X3]:((v3_setfam_1(X3,X1)&m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1))))=>v3_setfam_1(k4_subset_1(k1_zfmisc_1(X1),X2,X3),X1))))), inference(assume_negation,[status(cth)],[t63_setfam_1])).
fof(c_0_19, plain, ![X17, X18]:((~v3_setfam_1(X18,X17)|~r2_hidden(X17,X18)|~m1_subset_1(X18,k1_zfmisc_1(k1_zfmisc_1(X17))))&(r2_hidden(X17,X18)|v3_setfam_1(X18,X17)|~m1_subset_1(X18,k1_zfmisc_1(k1_zfmisc_1(X17))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_setfam_1])])])])).
fof(c_0_20, plain, ![X19, X20, X21]:(~m1_subset_1(X20,k1_zfmisc_1(X19))|~m1_subset_1(X21,k1_zfmisc_1(X19))|m1_subset_1(k4_subset_1(X19,X20,X21),k1_zfmisc_1(X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_subset_1])])).
fof(c_0_21, plain, ![X22, X23, X24]:(((~r2_hidden(X22,X23)|~r2_hidden(X22,X24)|~r2_hidden(X22,k5_xboole_0(X23,X24)))&(r2_hidden(X22,X23)|r2_hidden(X22,X24)|~r2_hidden(X22,k5_xboole_0(X23,X24))))&((~r2_hidden(X22,X23)|r2_hidden(X22,X24)|r2_hidden(X22,k5_xboole_0(X23,X24)))&(~r2_hidden(X22,X24)|r2_hidden(X22,X23)|r2_hidden(X22,k5_xboole_0(X23,X24))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_0])])])).
cnf(c_0_22, plain, (k4_subset_1(X2,X1,X3)=k3_tarski(k1_enumset1(X1,X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_23, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[c_0_17, c_0_16])).
fof(c_0_24, negated_conjecture, (~v1_xboole_0(esk1_0)&((v3_setfam_1(esk2_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0))))&((v3_setfam_1(esk3_0,esk1_0)&m1_subset_1(esk3_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0))))&~v3_setfam_1(k4_subset_1(k1_zfmisc_1(esk1_0),esk2_0,esk3_0),esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_18])])])])).
cnf(c_0_25, plain, (r2_hidden(X1,X2)|v3_setfam_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (m1_subset_1(k4_subset_1(X2,X1,X3),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_27, plain, ![X25, X26, X27, X28, X29, X30, X31, X32]:((((r2_hidden(X28,X25)|~r2_hidden(X28,X27)|X27!=k4_xboole_0(X25,X26))&(~r2_hidden(X28,X26)|~r2_hidden(X28,X27)|X27!=k4_xboole_0(X25,X26)))&(~r2_hidden(X29,X25)|r2_hidden(X29,X26)|r2_hidden(X29,X27)|X27!=k4_xboole_0(X25,X26)))&((~r2_hidden(esk4_3(X30,X31,X32),X32)|(~r2_hidden(esk4_3(X30,X31,X32),X30)|r2_hidden(esk4_3(X30,X31,X32),X31))|X32=k4_xboole_0(X30,X31))&((r2_hidden(esk4_3(X30,X31,X32),X30)|r2_hidden(esk4_3(X30,X31,X32),X32)|X32=k4_xboole_0(X30,X31))&(~r2_hidden(esk4_3(X30,X31,X32),X31)|r2_hidden(esk4_3(X30,X31,X32),X32)|X32=k4_xboole_0(X30,X31))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_28, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, plain, (k4_subset_1(X1,X2,X3)=k5_xboole_0(X2,k4_xboole_0(X3,X2))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_30, negated_conjecture, (~v3_setfam_1(k4_subset_1(k1_zfmisc_1(esk1_0),esk2_0,esk3_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, plain, (r2_hidden(X1,k4_subset_1(k1_zfmisc_1(X1),X2,X3))|v3_setfam_1(k4_subset_1(k1_zfmisc_1(X1),X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_32, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, plain, (~v3_setfam_1(X1,X2)|~r2_hidden(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_35, negated_conjecture, (v3_setfam_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_36, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,k4_subset_1(X4,X3,X2))|~m1_subset_1(X2,k1_zfmisc_1(X4))|~m1_subset_1(X3,k1_zfmisc_1(X4))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_38, negated_conjecture, (r2_hidden(esk1_0,k4_subset_1(k1_zfmisc_1(esk1_0),esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32]), c_0_33])])).
cnf(c_0_39, negated_conjecture, (~r2_hidden(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_33]), c_0_35])])).
cnf(c_0_40, negated_conjecture, (v3_setfam_1(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_41, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_36])).
cnf(c_0_42, negated_conjecture, (r2_hidden(esk1_0,k4_xboole_0(esk3_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_32]), c_0_33])]), c_0_39])).
cnf(c_0_43, negated_conjecture, (~r2_hidden(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_32]), c_0_40])])).
cnf(c_0_44, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 45
# Proof object clause steps            : 26
# Proof object formula steps           : 19
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 21
# Processed clauses                    : 271
# ...of these trivial                  : 1
# ...subsumed                          : 158
# ...remaining for further processing  : 112
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 0
# Generated clauses                    : 1104
# ...of the previous two non-trivial   : 945
# Contextual simplify-reflections      : 4
# Paramodulations                      : 1081
# Factorizations                       : 20
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 84
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 65
# Current number of unprocessed clauses: 711
# ...number of literals in the above   : 2489
# Current number of archived formulas  : 0
# Current number of archived clauses   : 27
# Clause-clause subsumption calls (NU) : 1377
# Rec. Clause-clause subsumption calls : 808
# Non-unit clause-clause subsumptions  : 113
# Unit Clause-clause subsumption calls : 50
# Rewrite failures with RHS unbound    : 34
# BW rewrite match attempts            : 32
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 17546

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.003 s
# Total time               : 0.019 s
# Maximum resident set size: 28744 pages
