# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/6a8b90ff886d4eac8bceb511a1f4c0ee.E_input', d10_xboole_0)).
fof(t4_tsep_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X1)=>(r1_tarski(u1_struct_0(X2),u1_struct_0(X3))<=>m1_pre_topc(X2,X3))))), file('data/MPTP/proofs/6a8b90ff886d4eac8bceb511a1f4c0ee.E_input', t4_tsep_1)).
fof(t96_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,u1_struct_0(X2))=>X4=k1_funct_1(X3,X4)))=>r2_funct_2(u1_struct_0(X2),u1_struct_0(X1),k4_tmap_1(X1,X2),X3))))), file('data/MPTP/proofs/6a8b90ff886d4eac8bceb511a1f4c0ee.E_input', t96_tmap_1)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/6a8b90ff886d4eac8bceb511a1f4c0ee.E_input', dt_m1_pre_topc)).
fof(cc1_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>v2_pre_topc(X2))), file('data/MPTP/proofs/6a8b90ff886d4eac8bceb511a1f4c0ee.E_input', cc1_pre_topc)).
fof(t59_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X2))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X1))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))))=>(![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>(r2_hidden(X6,u1_struct_0(X3))=>k3_funct_2(u1_struct_0(X2),u1_struct_0(X1),X4,X6)=k1_funct_1(X5,X6)))=>r2_funct_2(u1_struct_0(X3),u1_struct_0(X1),k2_tmap_1(X2,X1,X4,X3),X5))))))), file('data/MPTP/proofs/6a8b90ff886d4eac8bceb511a1f4c0ee.E_input', t59_tmap_1)).
fof(dt_k4_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_pre_topc(X2,X1))=>(v1_funct_1(k4_tmap_1(X1,X2))&(v1_funct_2(k4_tmap_1(X1,X2),u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(k4_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))), file('data/MPTP/proofs/6a8b90ff886d4eac8bceb511a1f4c0ee.E_input', dt_k4_tmap_1)).
fof(d4_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:(m1_pre_topc(X4,X1)=>k2_tmap_1(X1,X2,X3,X4)=k2_partfun1(u1_struct_0(X1),u1_struct_0(X2),X3,u1_struct_0(X4)))))), file('data/MPTP/proofs/6a8b90ff886d4eac8bceb511a1f4c0ee.E_input', d4_tmap_1)).
fof(redefinition_k3_funct_2, axiom, ![X1, X2, X3, X4]:((~(v1_xboole_0(X1))&((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&m1_subset_1(X4,X1)))=>k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4)), file('data/MPTP/proofs/6a8b90ff886d4eac8bceb511a1f4c0ee.E_input', redefinition_k3_funct_2)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/6a8b90ff886d4eac8bceb511a1f4c0ee.E_input', d1_xboole_0)).
fof(t55_pre_topc, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>m1_subset_1(X3,u1_struct_0(X1))))), file('data/MPTP/proofs/6a8b90ff886d4eac8bceb511a1f4c0ee.E_input', t55_pre_topc)).
fof(t95_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,u1_struct_0(X2))=>k1_funct_1(k4_tmap_1(X1,X2),X3)=X3)))), file('data/MPTP/proofs/6a8b90ff886d4eac8bceb511a1f4c0ee.E_input', t95_tmap_1)).
fof(t173_funct_2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X4]:((~(v1_xboole_0(X4))&m1_subset_1(X4,k1_zfmisc_1(X1)))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,X4,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))))=>(![X6]:(m1_subset_1(X6,X1)=>(r2_hidden(X6,X4)=>k3_funct_2(X1,X2,X3,X6)=k1_funct_1(X5,X6)))=>k2_partfun1(X1,X2,X3,X4)=X5)))))), file('data/MPTP/proofs/6a8b90ff886d4eac8bceb511a1f4c0ee.E_input', t173_funct_2)).
fof(t1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/6a8b90ff886d4eac8bceb511a1f4c0ee.E_input', t1_tsep_1)).
fof(redefinition_r2_funct_2, axiom, ![X1, X2, X3, X4]:(((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&(v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))=>(r2_funct_2(X1,X2,X3,X4)<=>X3=X4)), file('data/MPTP/proofs/6a8b90ff886d4eac8bceb511a1f4c0ee.E_input', redefinition_r2_funct_2)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/6a8b90ff886d4eac8bceb511a1f4c0ee.E_input', fc2_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/6a8b90ff886d4eac8bceb511a1f4c0ee.E_input', dt_l1_pre_topc)).
fof(c_0_17, plain, ![X11, X12]:(((r1_tarski(X11,X12)|X11!=X12)&(r1_tarski(X12,X11)|X11!=X12))&(~r1_tarski(X11,X12)|~r1_tarski(X12,X11)|X11=X12)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_18, plain, ![X34, X35, X36]:((~r1_tarski(u1_struct_0(X35),u1_struct_0(X36))|m1_pre_topc(X35,X36)|~m1_pre_topc(X36,X34)|~m1_pre_topc(X35,X34)|(~v2_pre_topc(X34)|~l1_pre_topc(X34)))&(~m1_pre_topc(X35,X36)|r1_tarski(u1_struct_0(X35),u1_struct_0(X36))|~m1_pre_topc(X36,X34)|~m1_pre_topc(X35,X34)|(~v2_pre_topc(X34)|~l1_pre_topc(X34)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_tsep_1])])])])).
cnf(c_0_19, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_20, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,u1_struct_0(X2))=>X4=k1_funct_1(X3,X4)))=>r2_funct_2(u1_struct_0(X2),u1_struct_0(X1),k4_tmap_1(X1,X2),X3)))))), inference(assume_negation,[status(cth)],[t96_tmap_1])).
cnf(c_0_21, plain, (m1_pre_topc(X1,X2)|~r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_19])).
fof(c_0_23, negated_conjecture, ![X10]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,u1_struct_0(esk2_0),u1_struct_0(esk1_0))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))))&((~m1_subset_1(X10,u1_struct_0(esk1_0))|(~r2_hidden(X10,u1_struct_0(esk2_0))|X10=k1_funct_1(esk3_0,X10)))&~r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_20])])])])])).
fof(c_0_24, plain, ![X55, X56]:(~l1_pre_topc(X55)|(~m1_pre_topc(X56,X55)|l1_pre_topc(X56))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
fof(c_0_25, plain, ![X19, X20]:(~v2_pre_topc(X19)|~l1_pre_topc(X19)|(~m1_pre_topc(X20,X19)|v2_pre_topc(X20))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_pre_topc])])])).
fof(c_0_26, plain, ![X13, X14, X15, X16, X17]:((m1_subset_1(esk4_5(X13,X14,X15,X16,X17),u1_struct_0(X14))|r2_funct_2(u1_struct_0(X15),u1_struct_0(X13),k2_tmap_1(X14,X13,X16,X15),X17)|(~v1_funct_1(X17)|(~v1_funct_2(X17,u1_struct_0(X15),u1_struct_0(X13))|~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X15),u1_struct_0(X13))))))|(~v1_funct_1(X16)|(~v1_funct_2(X16,u1_struct_0(X14),u1_struct_0(X13))|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X14),u1_struct_0(X13))))))|(v2_struct_0(X15)|~m1_pre_topc(X15,X14))|(v2_struct_0(X14)|(~v2_pre_topc(X14)|~l1_pre_topc(X14)))|(v2_struct_0(X13)|(~v2_pre_topc(X13)|~l1_pre_topc(X13))))&((r2_hidden(esk4_5(X13,X14,X15,X16,X17),u1_struct_0(X15))|r2_funct_2(u1_struct_0(X15),u1_struct_0(X13),k2_tmap_1(X14,X13,X16,X15),X17)|(~v1_funct_1(X17)|(~v1_funct_2(X17,u1_struct_0(X15),u1_struct_0(X13))|~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X15),u1_struct_0(X13))))))|(~v1_funct_1(X16)|(~v1_funct_2(X16,u1_struct_0(X14),u1_struct_0(X13))|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X14),u1_struct_0(X13))))))|(v2_struct_0(X15)|~m1_pre_topc(X15,X14))|(v2_struct_0(X14)|(~v2_pre_topc(X14)|~l1_pre_topc(X14)))|(v2_struct_0(X13)|(~v2_pre_topc(X13)|~l1_pre_topc(X13))))&(k3_funct_2(u1_struct_0(X14),u1_struct_0(X13),X16,esk4_5(X13,X14,X15,X16,X17))!=k1_funct_1(X17,esk4_5(X13,X14,X15,X16,X17))|r2_funct_2(u1_struct_0(X15),u1_struct_0(X13),k2_tmap_1(X14,X13,X16,X15),X17)|(~v1_funct_1(X17)|(~v1_funct_2(X17,u1_struct_0(X15),u1_struct_0(X13))|~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X15),u1_struct_0(X13))))))|(~v1_funct_1(X16)|(~v1_funct_2(X16,u1_struct_0(X14),u1_struct_0(X13))|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X14),u1_struct_0(X13))))))|(v2_struct_0(X15)|~m1_pre_topc(X15,X14))|(v2_struct_0(X14)|(~v2_pre_topc(X14)|~l1_pre_topc(X14)))|(v2_struct_0(X13)|(~v2_pre_topc(X13)|~l1_pre_topc(X13)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t59_tmap_1])])])])])])).
cnf(c_0_27, plain, (m1_pre_topc(X1,X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, plain, (v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_33, plain, ![X27, X28]:((v1_funct_1(k4_tmap_1(X27,X28))|(v2_struct_0(X27)|(~v2_pre_topc(X27)|~l1_pre_topc(X27))|~m1_pre_topc(X28,X27)))&((v1_funct_2(k4_tmap_1(X27,X28),u1_struct_0(X28),u1_struct_0(X27))|(v2_struct_0(X27)|(~v2_pre_topc(X27)|~l1_pre_topc(X27))|~m1_pre_topc(X28,X27)))&(m1_subset_1(k4_tmap_1(X27,X28),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X28),u1_struct_0(X27))))|(v2_struct_0(X27)|(~v2_pre_topc(X27)|~l1_pre_topc(X27))|~m1_pre_topc(X28,X27))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_tmap_1])])])])).
fof(c_0_34, plain, ![X29, X30, X31, X32]:(v2_struct_0(X29)|(~v2_pre_topc(X29)|~l1_pre_topc(X29))|(v2_struct_0(X30)|(~v2_pre_topc(X30)|~l1_pre_topc(X30))|(~v1_funct_1(X31)|(~v1_funct_2(X31,u1_struct_0(X29),u1_struct_0(X30))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X29),u1_struct_0(X30)))))|(~m1_pre_topc(X32,X29)|k2_tmap_1(X29,X30,X31,X32)=k2_partfun1(u1_struct_0(X29),u1_struct_0(X30),X31,u1_struct_0(X32)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_tmap_1])])])])).
cnf(c_0_35, plain, (r2_hidden(esk4_5(X1,X2,X3,X4,X5),u1_struct_0(X3))|r2_funct_2(u1_struct_0(X3),u1_struct_0(X1),k2_tmap_1(X2,X1,X4,X3),X5)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X1))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_1(X4)|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_pre_topc(X3,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, negated_conjecture, (m1_pre_topc(esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30])])).
cnf(c_0_37, negated_conjecture, (l1_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_28]), c_0_29])])).
cnf(c_0_38, negated_conjecture, (v2_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_28]), c_0_29]), c_0_30])])).
cnf(c_0_39, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_40, plain, (v1_funct_2(k4_tmap_1(X1,X2),u1_struct_0(X2),u1_struct_0(X1))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_41, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_42, plain, (m1_subset_1(k4_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_43, plain, (v1_funct_1(k4_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_44, plain, (v2_struct_0(X1)|v2_struct_0(X2)|k2_tmap_1(X1,X2,X3,X4)=k2_partfun1(u1_struct_0(X1),u1_struct_0(X2),X3,u1_struct_0(X4))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_pre_topc(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45, plain, (m1_subset_1(esk4_5(X1,X2,X3,X4,X5),u1_struct_0(X2))|r2_funct_2(u1_struct_0(X3),u1_struct_0(X1),k2_tmap_1(X2,X1,X4,X3),X5)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X1))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_1(X4)|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_pre_topc(X3,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_46, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(X1),k2_tmap_1(esk2_0,X1,X2,esk2_0),X3)|r2_hidden(esk4_5(X1,esk2_0,esk2_0,X2,X3),u1_struct_0(esk2_0))|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(esk2_0),u1_struct_0(X1))|~v1_funct_2(X2,u1_struct_0(esk2_0),u1_struct_0(X1))|~v1_funct_1(X3)|~v1_funct_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_47, negated_conjecture, (v1_funct_2(k4_tmap_1(esk1_0,esk2_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_28]), c_0_29]), c_0_30])]), c_0_41])).
cnf(c_0_48, negated_conjecture, (m1_subset_1(k4_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_28]), c_0_29]), c_0_30])]), c_0_41])).
cnf(c_0_49, negated_conjecture, (v1_funct_1(k4_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_28]), c_0_29]), c_0_30])]), c_0_41])).
cnf(c_0_50, negated_conjecture, (k2_tmap_1(esk2_0,X1,X2,esk2_0)=k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(X1),X2,u1_struct_0(esk2_0))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(esk2_0),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_36]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_51, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk2_0),u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_52, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_53, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_54, plain, ![X47, X48, X49, X50]:(v1_xboole_0(X47)|(~v1_funct_1(X49)|(~v1_funct_2(X49,X47,X48)|~m1_subset_1(X49,k1_zfmisc_1(k2_zfmisc_1(X47,X48))))|~m1_subset_1(X50,X47))|k3_funct_2(X47,X48,X49,X50)=k1_funct_1(X49,X50)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_funct_2])])])).
cnf(c_0_55, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(X1),k2_tmap_1(esk2_0,X1,X2,esk2_0),X3)|m1_subset_1(esk4_5(X1,esk2_0,esk2_0,X2,X3),u1_struct_0(esk2_0))|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(esk2_0),u1_struct_0(X1))|~v1_funct_2(X2,u1_struct_0(esk2_0),u1_struct_0(X1))|~v1_funct_1(X3)|~v1_funct_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_36]), c_0_37]), c_0_38])]), c_0_39])).
fof(c_0_56, plain, ![X51, X52, X53]:((~v1_xboole_0(X51)|~r2_hidden(X52,X51))&(r2_hidden(esk6_1(X53),X53)|v1_xboole_0(X53))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_57, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tmap_1(esk2_0,esk1_0,X1,esk2_0),k4_tmap_1(esk1_0,esk2_0))|r2_hidden(esk4_5(esk1_0,esk2_0,esk2_0,X1,k4_tmap_1(esk1_0,esk2_0)),u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48]), c_0_49]), c_0_29]), c_0_30])]), c_0_41])).
cnf(c_0_58, negated_conjecture, (k2_tmap_1(esk2_0,esk1_0,esk3_0,esk2_0)=k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52]), c_0_53]), c_0_29]), c_0_30])]), c_0_41])).
fof(c_0_59, plain, ![X24, X25, X26]:(v2_struct_0(X24)|~l1_pre_topc(X24)|(v2_struct_0(X25)|~m1_pre_topc(X25,X24)|(~m1_subset_1(X26,u1_struct_0(X25))|m1_subset_1(X26,u1_struct_0(X24))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t55_pre_topc])])])])).
cnf(c_0_60, plain, (v1_xboole_0(X1)|k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|~v1_funct_1(X2)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_61, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tmap_1(esk2_0,esk1_0,X1,esk2_0),k4_tmap_1(esk1_0,esk2_0))|m1_subset_1(esk4_5(esk1_0,esk2_0,esk2_0,X1,k4_tmap_1(esk1_0,esk2_0)),u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_47]), c_0_48]), c_0_49]), c_0_29]), c_0_30])]), c_0_41])).
cnf(c_0_62, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_63, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0)),k4_tmap_1(esk1_0,esk2_0))|r2_hidden(esk4_5(esk1_0,esk2_0,esk2_0,esk3_0,k4_tmap_1(esk1_0,esk2_0)),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_51]), c_0_58]), c_0_52]), c_0_53])])).
fof(c_0_64, plain, ![X21, X22, X23]:(v2_struct_0(X21)|(~v2_pre_topc(X21)|~l1_pre_topc(X21))|(v2_struct_0(X22)|~m1_pre_topc(X22,X21)|(~m1_subset_1(X23,u1_struct_0(X21))|(~r2_hidden(X23,u1_struct_0(X22))|k1_funct_1(k4_tmap_1(X21,X22),X23)=X23)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t95_tmap_1])])])])).
cnf(c_0_65, plain, (v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(X3,u1_struct_0(X1))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_59])).
fof(c_0_66, plain, ![X37, X38, X39, X40, X41]:((m1_subset_1(esk5_5(X37,X38,X39,X40,X41),X37)|k2_partfun1(X37,X38,X39,X40)=X41|(~v1_funct_1(X41)|(~v1_funct_2(X41,X40,X38)|~m1_subset_1(X41,k1_zfmisc_1(k2_zfmisc_1(X40,X38)))))|(v1_xboole_0(X40)|~m1_subset_1(X40,k1_zfmisc_1(X37)))|(~v1_funct_1(X39)|(~v1_funct_2(X39,X37,X38)|~m1_subset_1(X39,k1_zfmisc_1(k2_zfmisc_1(X37,X38)))))|v1_xboole_0(X38)|v1_xboole_0(X37))&((r2_hidden(esk5_5(X37,X38,X39,X40,X41),X40)|k2_partfun1(X37,X38,X39,X40)=X41|(~v1_funct_1(X41)|(~v1_funct_2(X41,X40,X38)|~m1_subset_1(X41,k1_zfmisc_1(k2_zfmisc_1(X40,X38)))))|(v1_xboole_0(X40)|~m1_subset_1(X40,k1_zfmisc_1(X37)))|(~v1_funct_1(X39)|(~v1_funct_2(X39,X37,X38)|~m1_subset_1(X39,k1_zfmisc_1(k2_zfmisc_1(X37,X38)))))|v1_xboole_0(X38)|v1_xboole_0(X37))&(k3_funct_2(X37,X38,X39,esk5_5(X37,X38,X39,X40,X41))!=k1_funct_1(X41,esk5_5(X37,X38,X39,X40,X41))|k2_partfun1(X37,X38,X39,X40)=X41|(~v1_funct_1(X41)|(~v1_funct_2(X41,X40,X38)|~m1_subset_1(X41,k1_zfmisc_1(k2_zfmisc_1(X40,X38)))))|(v1_xboole_0(X40)|~m1_subset_1(X40,k1_zfmisc_1(X37)))|(~v1_funct_1(X39)|(~v1_funct_2(X39,X37,X38)|~m1_subset_1(X39,k1_zfmisc_1(k2_zfmisc_1(X37,X38)))))|v1_xboole_0(X38)|v1_xboole_0(X37)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t173_funct_2])])])])])])).
fof(c_0_67, plain, ![X57, X58]:(~l1_pre_topc(X57)|(~m1_pre_topc(X58,X57)|m1_subset_1(u1_struct_0(X58),k1_zfmisc_1(u1_struct_0(X57))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])).
cnf(c_0_68, negated_conjecture, (k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,X1)=k1_funct_1(esk3_0,X1)|v1_xboole_0(u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_51]), c_0_52]), c_0_53])])).
cnf(c_0_69, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0)),k4_tmap_1(esk1_0,esk2_0))|m1_subset_1(esk4_5(esk1_0,esk2_0,esk2_0,esk3_0,k4_tmap_1(esk1_0,esk2_0)),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_51]), c_0_58]), c_0_52]), c_0_53])])).
cnf(c_0_70, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0)),k4_tmap_1(esk1_0,esk2_0))|~v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_71, plain, (v2_struct_0(X1)|v2_struct_0(X2)|k1_funct_1(k4_tmap_1(X1,X2),X3)=X3|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_72, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_28]), c_0_29])]), c_0_39]), c_0_41])).
cnf(c_0_73, plain, (m1_subset_1(esk5_5(X1,X2,X3,X4,X5),X1)|k2_partfun1(X1,X2,X3,X4)=X5|v1_xboole_0(X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_funct_1(X5)|~v1_funct_2(X5,X4,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_74, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_75, plain, (r2_funct_2(u1_struct_0(X4),u1_struct_0(X2),k2_tmap_1(X1,X2,X3,X4),X5)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X2)|k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,esk4_5(X2,X1,X4,X3,X5))!=k1_funct_1(X5,esk4_5(X2,X1,X4,X3,X5))|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_pre_topc(X4,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_76, negated_conjecture, (k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk4_5(esk1_0,esk2_0,esk2_0,esk3_0,k4_tmap_1(esk1_0,esk2_0)))=k1_funct_1(esk3_0,esk4_5(esk1_0,esk2_0,esk2_0,esk3_0,k4_tmap_1(esk1_0,esk2_0)))|r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0)),k4_tmap_1(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70])).
cnf(c_0_77, negated_conjecture, (k1_funct_1(k4_tmap_1(X1,esk2_0),esk4_5(esk1_0,esk2_0,esk2_0,esk3_0,k4_tmap_1(esk1_0,esk2_0)))=esk4_5(esk1_0,esk2_0,esk2_0,esk3_0,k4_tmap_1(esk1_0,esk2_0))|r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0)),k4_tmap_1(esk1_0,esk2_0))|v2_struct_0(X1)|~m1_subset_1(esk4_5(esk1_0,esk2_0,esk2_0,esk3_0,k4_tmap_1(esk1_0,esk2_0)),u1_struct_0(X1))|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_63]), c_0_39])).
cnf(c_0_78, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0)),k4_tmap_1(esk1_0,esk2_0))|m1_subset_1(esk4_5(esk1_0,esk2_0,esk2_0,esk3_0,k4_tmap_1(esk1_0,esk2_0)),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_72, c_0_69])).
cnf(c_0_79, negated_conjecture, (X1=k1_funct_1(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_80, negated_conjecture, (k2_partfun1(X1,u1_struct_0(esk1_0),X2,u1_struct_0(esk2_0))=esk3_0|v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))|v1_xboole_0(X1)|m1_subset_1(esk5_5(X1,u1_struct_0(esk1_0),X2,u1_struct_0(esk2_0),esk3_0),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(esk1_0))))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(X1))|~v1_funct_2(X2,X1,u1_struct_0(esk1_0))|~v1_funct_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_51]), c_0_52]), c_0_53])])).
cnf(c_0_81, negated_conjecture, (m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_36]), c_0_37])])).
fof(c_0_82, plain, ![X43, X44, X45, X46]:((~r2_funct_2(X43,X44,X45,X46)|X45=X46|(~v1_funct_1(X45)|(~v1_funct_2(X45,X43,X44)|~m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(X43,X44))))|(~v1_funct_1(X46)|(~v1_funct_2(X46,X43,X44)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(X43,X44)))))))&(X45!=X46|r2_funct_2(X43,X44,X45,X46)|(~v1_funct_1(X45)|(~v1_funct_2(X45,X43,X44)|~m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(X43,X44))))|(~v1_funct_1(X46)|(~v1_funct_2(X46,X43,X44)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(X43,X44)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r2_funct_2])])])).
cnf(c_0_83, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0)),k4_tmap_1(esk1_0,esk2_0))|k1_funct_1(k4_tmap_1(esk1_0,esk2_0),esk4_5(esk1_0,esk2_0,esk2_0,esk3_0,k4_tmap_1(esk1_0,esk2_0)))!=k1_funct_1(esk3_0,esk4_5(esk1_0,esk2_0,esk2_0,esk3_0,k4_tmap_1(esk1_0,esk2_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_58]), c_0_48]), c_0_52]), c_0_47]), c_0_51]), c_0_49]), c_0_53]), c_0_36]), c_0_29]), c_0_37]), c_0_30]), c_0_38])]), c_0_39]), c_0_41])).
cnf(c_0_84, negated_conjecture, (k1_funct_1(k4_tmap_1(esk1_0,esk2_0),esk4_5(esk1_0,esk2_0,esk2_0,esk3_0,k4_tmap_1(esk1_0,esk2_0)))=esk4_5(esk1_0,esk2_0,esk2_0,esk3_0,k4_tmap_1(esk1_0,esk2_0))|r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0)),k4_tmap_1(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_28]), c_0_29]), c_0_30])]), c_0_41]), c_0_78])).
cnf(c_0_85, negated_conjecture, (k1_funct_1(esk3_0,esk4_5(esk1_0,esk2_0,esk2_0,esk3_0,k4_tmap_1(esk1_0,esk2_0)))=esk4_5(esk1_0,esk2_0,esk2_0,esk3_0,k4_tmap_1(esk1_0,esk2_0))|r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0)),k4_tmap_1(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_63]), c_0_78])).
cnf(c_0_86, negated_conjecture, (k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0))=esk3_0|v1_xboole_0(u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))|m1_subset_1(esk5_5(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0),esk3_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_51]), c_0_52]), c_0_81]), c_0_53])])).
cnf(c_0_87, plain, (X3=X4|~r2_funct_2(X1,X2,X3,X4)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X4)|~v1_funct_2(X4,X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_88, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0)),k4_tmap_1(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_85])).
cnf(c_0_89, plain, (k2_partfun1(X1,X2,X3,X4)=X5|v1_xboole_0(X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|k3_funct_2(X1,X2,X3,esk5_5(X1,X2,X3,X4,X5))!=k1_funct_1(X5,esk5_5(X1,X2,X3,X4,X5))|~v1_funct_1(X5)|~v1_funct_2(X5,X4,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_90, negated_conjecture, (k3_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk5_5(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0),esk3_0))=k1_funct_1(esk3_0,esk5_5(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0),esk3_0))|k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0))=esk3_0|v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_68, c_0_86])).
cnf(c_0_91, plain, (r2_funct_2(X3,X4,X1,X2)|X1!=X2|~v1_funct_1(X1)|~v1_funct_2(X1,X3,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X2)|~v1_funct_2(X2,X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_92, negated_conjecture, (k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0))=k4_tmap_1(esk1_0,esk2_0)|~m1_subset_1(k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0)),u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_48]), c_0_47]), c_0_49])])).
cnf(c_0_93, negated_conjecture, (k2_partfun1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,u1_struct_0(esk2_0))=esk3_0|v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_52]), c_0_81]), c_0_51]), c_0_53])])).
cnf(c_0_94, plain, (r2_funct_2(X1,X2,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)), inference(er,[status(thm)],[c_0_91])).
fof(c_0_95, plain, ![X59]:(v2_struct_0(X59)|~l1_struct_0(X59)|~v1_xboole_0(u1_struct_0(X59))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_96, negated_conjecture, (~r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_97, negated_conjecture, (k4_tmap_1(esk1_0,esk2_0)=esk3_0|v1_xboole_0(u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_52]), c_0_51]), c_0_53])])).
cnf(c_0_98, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_51]), c_0_52]), c_0_53])])).
cnf(c_0_99, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_100, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_98])])).
fof(c_0_101, plain, ![X33]:(~l1_pre_topc(X33)|l1_struct_0(X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_102, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~l1_struct_0(esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_39])).
cnf(c_0_103, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_104, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_37])])).
cnf(c_0_105, negated_conjecture, (~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_104]), c_0_41])).
cnf(c_0_106, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_103]), c_0_29])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 107
# Proof object clause steps            : 72
# Proof object formula steps           : 35
# Proof object conjectures             : 50
# Proof object clause conjectures      : 47
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 32
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 38
# Proof object simplifying inferences  : 124
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 37
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 37
# Processed clauses                    : 206
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 205
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 38
# Backward-rewritten                   : 17
# Generated clauses                    : 210
# ...of the previous two non-trivial   : 181
# Contextual simplify-reflections      : 19
# Paramodulations                      : 207
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 111
#    Positive orientable unit clauses  : 27
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 80
# Current number of unprocessed clauses: 38
# ...number of literals in the above   : 152
# Current number of archived formulas  : 0
# Current number of archived clauses   : 91
# Clause-clause subsumption calls (NU) : 5985
# Rec. Clause-clause subsumption calls : 303
# Non-unit clause-clause subsumptions  : 57
# Unit Clause-clause subsumption calls : 132
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 30
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 15751

# -------------------------------------------------
# User time                : 0.018 s
# System time              : 0.003 s
# Total time               : 0.021 s
# Maximum resident set size: 28748 pages
