# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t48_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:((~(v1_xboole_0(X2))&(v1_waybel_0(X2,X1)&(v12_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v14_waybel_0(X2,X1)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&X2=k5_waybel_0(X1,X3))))), file('data/MPTP/proofs/6ad27625974f45559b041aa2ca24db85.E_input', t48_waybel_0)).
fof(fraenkel_a_2_0_waybel_0, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&l1_orders_2(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r2_hidden(X1,a_2_0_waybel_0(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X4,X5)&r2_hidden(X5,X3))))))), file('data/MPTP/proofs/6ad27625974f45559b041aa2ca24db85.E_input', fraenkel_a_2_0_waybel_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/6ad27625974f45559b041aa2ca24db85.E_input', t24_orders_2)).
fof(t31_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>r2_lattice3(X1,k3_waybel_0(X1,X2),X3))))), file('data/MPTP/proofs/6ad27625974f45559b041aa2ca24db85.E_input', t31_waybel_0)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/6ad27625974f45559b041aa2ca24db85.E_input', d9_lattice3)).
fof(t14_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k3_waybel_0(X1,X2)=a_2_0_waybel_0(X1,X2))), file('data/MPTP/proofs/6ad27625974f45559b041aa2ca24db85.E_input', t14_waybel_0)).
fof(t17_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2))))), file('data/MPTP/proofs/6ad27625974f45559b041aa2ca24db85.E_input', t17_waybel_0)).
fof(dt_k5_waybel_0, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/6ad27625974f45559b041aa2ca24db85.E_input', dt_k5_waybel_0)).
fof(d19_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v12_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r1_orders_2(X1,X4,X3))=>r2_hidden(X4,X2))))))), file('data/MPTP/proofs/6ad27625974f45559b041aa2ca24db85.E_input', d19_waybel_0)).
fof(d15_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k3_waybel_0(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>?[X5]:(m1_subset_1(X5,u1_struct_0(X1))&(r1_orders_2(X1,X4,X5)&r2_hidden(X5,X2))))))))), file('data/MPTP/proofs/6ad27625974f45559b041aa2ca24db85.E_input', d15_waybel_0)).
fof(d21_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:((~(v1_xboole_0(X2))&(v1_waybel_0(X2,X1)&(v12_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v14_waybel_0(X2,X1)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_hidden(X3,X2)&r2_lattice3(X1,X2,X3)))))), file('data/MPTP/proofs/6ad27625974f45559b041aa2ca24db85.E_input', d21_waybel_0)).
fof(c_0_11, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:((~(v1_xboole_0(X2))&(v1_waybel_0(X2,X1)&(v12_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v14_waybel_0(X2,X1)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&X2=k5_waybel_0(X1,X3)))))), inference(assume_negation,[status(cth)],[t48_waybel_0])).
fof(c_0_12, plain, ![X35, X36, X37, X40, X41]:(((m1_subset_1(esk10_3(X35,X36,X37),u1_struct_0(X36))|~r2_hidden(X35,a_2_0_waybel_0(X36,X37))|(v2_struct_0(X36)|~l1_orders_2(X36)|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X36)))))&((X35=esk10_3(X35,X36,X37)|~r2_hidden(X35,a_2_0_waybel_0(X36,X37))|(v2_struct_0(X36)|~l1_orders_2(X36)|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X36)))))&((m1_subset_1(esk11_3(X35,X36,X37),u1_struct_0(X36))|~r2_hidden(X35,a_2_0_waybel_0(X36,X37))|(v2_struct_0(X36)|~l1_orders_2(X36)|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X36)))))&((r1_orders_2(X36,esk10_3(X35,X36,X37),esk11_3(X35,X36,X37))|~r2_hidden(X35,a_2_0_waybel_0(X36,X37))|(v2_struct_0(X36)|~l1_orders_2(X36)|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X36)))))&(r2_hidden(esk11_3(X35,X36,X37),X37)|~r2_hidden(X35,a_2_0_waybel_0(X36,X37))|(v2_struct_0(X36)|~l1_orders_2(X36)|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X36)))))))))&(~m1_subset_1(X40,u1_struct_0(X36))|(X35!=X40|(~m1_subset_1(X41,u1_struct_0(X36))|(~r1_orders_2(X36,X40,X41)|~r2_hidden(X41,X37))))|r2_hidden(X35,a_2_0_waybel_0(X36,X37))|(v2_struct_0(X36)|~l1_orders_2(X36)|~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X36)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_0_waybel_0])])])])])])).
fof(c_0_13, plain, ![X28, X29]:(v2_struct_0(X28)|(~v3_orders_2(X28)|~l1_orders_2(X28))|(~m1_subset_1(X29,u1_struct_0(X28))|r1_orders_2(X28,X29,X29))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_14, negated_conjecture, ![X23]:((~v2_struct_0(esk6_0)&(v3_orders_2(esk6_0)&(v4_orders_2(esk6_0)&l1_orders_2(esk6_0))))&((~v1_xboole_0(esk7_0)&(v1_waybel_0(esk7_0,esk6_0)&(v12_waybel_0(esk7_0,esk6_0)&m1_subset_1(esk7_0,k1_zfmisc_1(u1_struct_0(esk6_0))))))&((~v14_waybel_0(esk7_0,esk6_0)|(~m1_subset_1(X23,u1_struct_0(esk6_0))|esk7_0!=k5_waybel_0(esk6_0,X23)))&((m1_subset_1(esk8_0,u1_struct_0(esk6_0))|v14_waybel_0(esk7_0,esk6_0))&(esk7_0=k5_waybel_0(esk6_0,esk8_0)|v14_waybel_0(esk7_0,esk6_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])])])).
fof(c_0_15, plain, ![X46, X47, X48]:((~r2_lattice3(X46,X47,X48)|r2_lattice3(X46,k3_waybel_0(X46,X47),X48)|~m1_subset_1(X48,u1_struct_0(X46))|~m1_subset_1(X47,k1_zfmisc_1(u1_struct_0(X46)))|(v2_struct_0(X46)|(~v3_orders_2(X46)|(~v4_orders_2(X46)|~l1_orders_2(X46)))))&(~r2_lattice3(X46,k3_waybel_0(X46,X47),X48)|r2_lattice3(X46,X47,X48)|~m1_subset_1(X48,u1_struct_0(X46))|~m1_subset_1(X47,k1_zfmisc_1(u1_struct_0(X46)))|(v2_struct_0(X46)|(~v3_orders_2(X46)|(~v4_orders_2(X46)|~l1_orders_2(X46)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t31_waybel_0])])])])])).
cnf(c_0_16, plain, (r2_hidden(X3,a_2_0_waybel_0(X2,X5))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~m1_subset_1(X4,u1_struct_0(X2))|~r1_orders_2(X2,X1,X4)|~r2_hidden(X4,X5)|~l1_orders_2(X2)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, negated_conjecture, (v3_orders_2(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, negated_conjecture, (l1_orders_2(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, negated_conjecture, (~v2_struct_0(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_21, plain, ![X30, X31, X32, X33]:((~r2_lattice3(X30,X31,X32)|(~m1_subset_1(X33,u1_struct_0(X30))|(~r2_hidden(X33,X31)|r1_orders_2(X30,X33,X32)))|~m1_subset_1(X32,u1_struct_0(X30))|~l1_orders_2(X30))&((m1_subset_1(esk9_3(X30,X31,X32),u1_struct_0(X30))|r2_lattice3(X30,X31,X32)|~m1_subset_1(X32,u1_struct_0(X30))|~l1_orders_2(X30))&((r2_hidden(esk9_3(X30,X31,X32),X31)|r2_lattice3(X30,X31,X32)|~m1_subset_1(X32,u1_struct_0(X30))|~l1_orders_2(X30))&(~r1_orders_2(X30,esk9_3(X30,X31,X32),X32)|r2_lattice3(X30,X31,X32)|~m1_subset_1(X32,u1_struct_0(X30))|~l1_orders_2(X30))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
cnf(c_0_22, plain, (r2_lattice3(X1,k3_waybel_0(X1,X2),X3)|v2_struct_0(X1)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, negated_conjecture, (v4_orders_2(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_24, plain, (v2_struct_0(X1)|r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~r1_orders_2(X1,X2,X4)|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_25, negated_conjecture, (r1_orders_2(esk6_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])]), c_0_20])).
fof(c_0_26, plain, ![X49, X50]:(v2_struct_0(X49)|~l1_orders_2(X49)|(~m1_subset_1(X50,k1_zfmisc_1(u1_struct_0(X49)))|k3_waybel_0(X49,X50)=a_2_0_waybel_0(X49,X50))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_waybel_0])])])])).
fof(c_0_27, plain, ![X25, X26, X27]:((~r2_hidden(X27,k5_waybel_0(X25,X26))|r1_orders_2(X25,X27,X26)|~m1_subset_1(X27,u1_struct_0(X25))|~m1_subset_1(X26,u1_struct_0(X25))|(v2_struct_0(X25)|~l1_orders_2(X25)))&(~r1_orders_2(X25,X27,X26)|r2_hidden(X27,k5_waybel_0(X25,X26))|~m1_subset_1(X27,u1_struct_0(X25))|~m1_subset_1(X26,u1_struct_0(X25))|(v2_struct_0(X25)|~l1_orders_2(X25)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_waybel_0])])])])])).
cnf(c_0_28, plain, (r1_orders_2(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, negated_conjecture, (r2_lattice3(esk6_0,k3_waybel_0(esk6_0,X1),X2)|~r2_lattice3(esk6_0,X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X2,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_30, negated_conjecture, (r2_hidden(X1,a_2_0_waybel_0(esk6_0,X2))|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_19])]), c_0_20])).
cnf(c_0_31, plain, (v2_struct_0(X1)|k3_waybel_0(X1,X2)=a_2_0_waybel_0(X1,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, plain, (r1_orders_2(X2,X1,X3)|v2_struct_0(X2)|~r2_hidden(X1,k5_waybel_0(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, plain, (r2_hidden(esk9_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_34, negated_conjecture, (r1_orders_2(esk6_0,X1,X2)|~r2_lattice3(esk6_0,X3,X2)|~r2_hidden(X1,k3_waybel_0(esk6_0,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))|~m1_subset_1(X2,u1_struct_0(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_19])])).
cnf(c_0_35, negated_conjecture, (r2_hidden(X1,k3_waybel_0(esk6_0,X2))|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_19])]), c_0_20])).
cnf(c_0_36, plain, (r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk9_3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_37, plain, (r2_lattice3(X1,k5_waybel_0(X2,X3),X4)|v2_struct_0(X2)|r1_orders_2(X2,esk9_3(X1,k5_waybel_0(X2,X3),X4),X3)|~m1_subset_1(esk9_3(X1,k5_waybel_0(X2,X3),X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_38, plain, (m1_subset_1(esk9_3(X1,X2,X3),u1_struct_0(X1))|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_39, plain, (r2_hidden(X2,k5_waybel_0(X1,X3))|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_40, plain, (r1_orders_2(X1,esk10_3(X2,X1,X3),esk11_3(X2,X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~l1_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_41, plain, (m1_subset_1(esk10_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_42, plain, (m1_subset_1(esk11_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_43, negated_conjecture, (r1_orders_2(esk6_0,X1,X2)|~r2_lattice3(esk6_0,X3,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))|~m1_subset_1(X2,u1_struct_0(esk6_0))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_44, plain, (r2_lattice3(X1,k5_waybel_0(X1,X2),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_45, plain, (v2_struct_0(X1)|r2_hidden(esk10_3(X2,X1,X3),k5_waybel_0(X1,esk11_3(X2,X1,X3)))|~r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41]), c_0_42])).
cnf(c_0_46, plain, (X1=esk10_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_47, negated_conjecture, (r1_orders_2(esk6_0,X1,X2)|~r2_hidden(X1,k5_waybel_0(esk6_0,X2))|~m1_subset_1(k5_waybel_0(esk6_0,X2),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))|~m1_subset_1(X2,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_19])]), c_0_20])).
cnf(c_0_48, plain, (v2_struct_0(X1)|r2_hidden(X2,k5_waybel_0(X1,esk11_3(X2,X1,X3)))|~r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
fof(c_0_49, plain, ![X51, X52]:(v2_struct_0(X51)|~l1_orders_2(X51)|~m1_subset_1(X52,u1_struct_0(X51))|m1_subset_1(k5_waybel_0(X51,X52),k1_zfmisc_1(u1_struct_0(X51)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_waybel_0])])])).
fof(c_0_50, plain, ![X6, X7, X8, X9]:((~v12_waybel_0(X7,X6)|(~m1_subset_1(X8,u1_struct_0(X6))|(~m1_subset_1(X9,u1_struct_0(X6))|(~r2_hidden(X8,X7)|~r1_orders_2(X6,X9,X8)|r2_hidden(X9,X7))))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))&((m1_subset_1(esk1_2(X6,X7),u1_struct_0(X6))|v12_waybel_0(X7,X6)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))&((m1_subset_1(esk2_2(X6,X7),u1_struct_0(X6))|v12_waybel_0(X7,X6)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))&(((r2_hidden(esk1_2(X6,X7),X7)|v12_waybel_0(X7,X6)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))&(r1_orders_2(X6,esk2_2(X6,X7),esk1_2(X6,X7))|v12_waybel_0(X7,X6)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6)))&(~r2_hidden(esk2_2(X6,X7),X7)|v12_waybel_0(X7,X6)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d19_waybel_0])])])])])).
cnf(c_0_51, negated_conjecture, (r1_orders_2(esk6_0,X1,esk11_3(X1,esk6_0,X2))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,X2))|~m1_subset_1(k5_waybel_0(esk6_0,esk11_3(X1,esk6_0,X2)),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(esk11_3(X1,esk6_0,X2),u1_struct_0(esk6_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_19])]), c_0_20])).
cnf(c_0_52, plain, (v2_struct_0(X1)|m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_53, plain, (r2_hidden(X4,X1)|~v12_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,X1)|~r1_orders_2(X2,X4,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_54, plain, (r1_orders_2(X1,esk2_2(X1,X2),esk1_2(X1,X2))|v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_55, plain, (m1_subset_1(esk1_2(X1,X2),u1_struct_0(X1))|v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_56, plain, (m1_subset_1(esk2_2(X1,X2),u1_struct_0(X1))|v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_57, negated_conjecture, (r1_orders_2(esk6_0,X1,esk11_3(X1,esk6_0,X2))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,X2))|~m1_subset_1(esk11_3(X1,esk6_0,X2),u1_struct_0(esk6_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_19])]), c_0_20])).
cnf(c_0_58, plain, (v12_waybel_0(X2,X1)|~r2_hidden(esk2_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_59, plain, (r2_hidden(esk2_2(X1,X2),X3)|v12_waybel_0(X2,X1)|~r2_hidden(esk1_2(X1,X2),X3)|~v12_waybel_0(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55]), c_0_56])).
cnf(c_0_60, negated_conjecture, (v12_waybel_0(esk7_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_61, negated_conjecture, (m1_subset_1(esk7_0,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_62, negated_conjecture, (r2_hidden(X1,X2)|~r2_hidden(esk11_3(X1,esk6_0,X3),X2)|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,X3))|~v12_waybel_0(X2,esk6_0)|~m1_subset_1(esk11_3(X1,esk6_0,X3),u1_struct_0(esk6_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_57]), c_0_19])])).
cnf(c_0_63, plain, (r2_hidden(esk11_3(X1,X2,X3),X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_64, negated_conjecture, (v12_waybel_0(a_2_0_waybel_0(esk6_0,X1),X2)|~r2_hidden(esk2_2(X2,a_2_0_waybel_0(esk6_0,X1)),X1)|~m1_subset_1(esk2_2(X2,a_2_0_waybel_0(esk6_0,X1)),u1_struct_0(esk6_0))|~m1_subset_1(a_2_0_waybel_0(esk6_0,X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk6_0)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_58, c_0_30])).
cnf(c_0_65, negated_conjecture, (r2_hidden(esk2_2(esk6_0,X1),esk7_0)|v12_waybel_0(X1,esk6_0)|~r2_hidden(esk1_2(esk6_0,X1),esk7_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61]), c_0_19])])).
cnf(c_0_66, negated_conjecture, (r2_hidden(X1,X2)|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,X2))|~v12_waybel_0(X2,esk6_0)|~m1_subset_1(esk11_3(X1,esk6_0,X2),u1_struct_0(esk6_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_19])]), c_0_20])).
cnf(c_0_67, negated_conjecture, (v12_waybel_0(a_2_0_waybel_0(esk6_0,esk7_0),esk6_0)|~r2_hidden(esk1_2(esk6_0,a_2_0_waybel_0(esk6_0,esk7_0)),esk7_0)|~m1_subset_1(esk2_2(esk6_0,a_2_0_waybel_0(esk6_0,esk7_0)),u1_struct_0(esk6_0))|~m1_subset_1(a_2_0_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_61]), c_0_19])])).
cnf(c_0_68, negated_conjecture, (r2_hidden(X1,X2)|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,X2))|~v12_waybel_0(X2,esk6_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_42]), c_0_19])]), c_0_20])).
cnf(c_0_69, plain, (r2_hidden(esk1_2(X1,X2),X2)|v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_70, negated_conjecture, (v12_waybel_0(a_2_0_waybel_0(esk6_0,esk7_0),esk6_0)|~r2_hidden(esk1_2(esk6_0,a_2_0_waybel_0(esk6_0,esk7_0)),esk7_0)|~m1_subset_1(a_2_0_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_56]), c_0_19])])).
cnf(c_0_71, negated_conjecture, (r2_hidden(esk1_2(X1,a_2_0_waybel_0(esk6_0,X2)),X2)|v12_waybel_0(a_2_0_waybel_0(esk6_0,X2),X1)|~v12_waybel_0(X2,esk6_0)|~m1_subset_1(esk1_2(X1,a_2_0_waybel_0(esk6_0,X2)),u1_struct_0(esk6_0))|~m1_subset_1(a_2_0_waybel_0(esk6_0,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_72, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,esk11_3(X2,X1,X3))|~r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_40, c_0_46])).
cnf(c_0_73, plain, (v2_struct_0(X1)|m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_41, c_0_46])).
cnf(c_0_74, negated_conjecture, (v12_waybel_0(a_2_0_waybel_0(esk6_0,esk7_0),esk6_0)|~m1_subset_1(esk1_2(esk6_0,a_2_0_waybel_0(esk6_0,esk7_0)),u1_struct_0(esk6_0))|~m1_subset_1(a_2_0_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_60]), c_0_61]), c_0_19])])).
cnf(c_0_75, plain, (v2_struct_0(X1)|r2_hidden(X2,X3)|~r2_hidden(esk11_3(X2,X1,X4),X3)|~r2_hidden(X2,a_2_0_waybel_0(X1,X4))|~v12_waybel_0(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_72]), c_0_73]), c_0_42])).
cnf(c_0_76, negated_conjecture, (v12_waybel_0(a_2_0_waybel_0(esk6_0,esk7_0),esk6_0)|~m1_subset_1(a_2_0_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_55]), c_0_19])])).
cnf(c_0_77, negated_conjecture, (v2_struct_0(X1)|r2_hidden(X2,k3_waybel_0(esk6_0,X3))|~r2_hidden(esk11_3(X2,X1,X4),X3)|~r2_hidden(X2,a_2_0_waybel_0(X1,X4))|~v12_waybel_0(k3_waybel_0(esk6_0,X3),X1)|~m1_subset_1(k3_waybel_0(esk6_0,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk11_3(X2,X1,X4),u1_struct_0(esk6_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_75, c_0_35])).
cnf(c_0_78, negated_conjecture, (v12_waybel_0(k3_waybel_0(esk6_0,esk7_0),esk6_0)|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_31]), c_0_61]), c_0_19])]), c_0_20])).
cnf(c_0_79, plain, (v2_struct_0(X1)|r2_hidden(esk2_2(X1,X2),a_2_0_waybel_0(X1,X3))|v12_waybel_0(X2,X1)|~r2_hidden(esk1_2(X1,X2),X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_54]), c_0_56]), c_0_55])).
cnf(c_0_80, plain, (v2_struct_0(X1)|r2_hidden(esk10_3(X2,X1,X3),X4)|~r2_hidden(esk11_3(X2,X1,X3),X4)|~r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~v12_waybel_0(X4,X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_40]), c_0_42]), c_0_41])).
cnf(c_0_81, negated_conjecture, (r2_hidden(X1,k3_waybel_0(esk6_0,esk7_0))|~r2_hidden(esk11_3(X1,esk6_0,X2),esk7_0)|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,X2))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(esk11_3(X1,esk6_0,X2),u1_struct_0(esk6_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_61]), c_0_19])]), c_0_20])).
cnf(c_0_82, plain, (v2_struct_0(X1)|v12_waybel_0(a_2_0_waybel_0(X1,X2),X1)|~r2_hidden(esk1_2(X1,a_2_0_waybel_0(X1,X2)),X2)|~m1_subset_1(a_2_0_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_58, c_0_79])).
cnf(c_0_83, plain, (v2_struct_0(X1)|r2_hidden(esk10_3(X2,X1,X3),X3)|~r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~v12_waybel_0(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_80, c_0_63])).
cnf(c_0_84, negated_conjecture, (r2_hidden(X1,k3_waybel_0(esk6_0,esk7_0))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(esk11_3(X1,esk6_0,esk7_0),u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_63]), c_0_61]), c_0_19])]), c_0_20])).
cnf(c_0_85, plain, (v2_struct_0(X1)|r2_hidden(esk10_3(X2,X1,X3),a_2_0_waybel_0(X1,X4))|~r2_hidden(esk11_3(X2,X1,X3),X4)|~r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_40]), c_0_41]), c_0_42])).
cnf(c_0_86, negated_conjecture, (v12_waybel_0(a_2_0_waybel_0(esk6_0,X1),esk6_0)|~v12_waybel_0(X1,esk6_0)|~m1_subset_1(esk1_2(esk6_0,a_2_0_waybel_0(esk6_0,X1)),u1_struct_0(esk6_0))|~m1_subset_1(a_2_0_waybel_0(esk6_0,X1),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_71]), c_0_19])]), c_0_20])).
cnf(c_0_87, plain, (v2_struct_0(X1)|r2_hidden(X2,X3)|~r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~v12_waybel_0(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_83, c_0_46])).
cnf(c_0_88, negated_conjecture, (r2_hidden(X1,k3_waybel_0(esk6_0,esk7_0))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_42]), c_0_61]), c_0_19])]), c_0_20])).
cnf(c_0_89, plain, (v2_struct_0(X1)|r2_hidden(esk10_3(X2,X1,X3),a_2_0_waybel_0(X1,X3))|~r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_85, c_0_63])).
fof(c_0_90, plain, ![X12, X13, X14, X15, X17, X19]:((((m1_subset_1(esk3_4(X12,X13,X14,X15),u1_struct_0(X12))|~r2_hidden(X15,X14)|~m1_subset_1(X15,u1_struct_0(X12))|X14!=k3_waybel_0(X12,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12))&((r1_orders_2(X12,X15,esk3_4(X12,X13,X14,X15))|~r2_hidden(X15,X14)|~m1_subset_1(X15,u1_struct_0(X12))|X14!=k3_waybel_0(X12,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12))&(r2_hidden(esk3_4(X12,X13,X14,X15),X13)|~r2_hidden(X15,X14)|~m1_subset_1(X15,u1_struct_0(X12))|X14!=k3_waybel_0(X12,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12))))&(~m1_subset_1(X17,u1_struct_0(X12))|(~r1_orders_2(X12,X15,X17)|~r2_hidden(X17,X13))|r2_hidden(X15,X14)|~m1_subset_1(X15,u1_struct_0(X12))|X14!=k3_waybel_0(X12,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12)))&((m1_subset_1(esk4_3(X12,X13,X14),u1_struct_0(X12))|X14=k3_waybel_0(X12,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12))&((~r2_hidden(esk4_3(X12,X13,X14),X14)|(~m1_subset_1(X19,u1_struct_0(X12))|(~r1_orders_2(X12,esk4_3(X12,X13,X14),X19)|~r2_hidden(X19,X13)))|X14=k3_waybel_0(X12,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12))&((m1_subset_1(esk5_3(X12,X13,X14),u1_struct_0(X12))|r2_hidden(esk4_3(X12,X13,X14),X14)|X14=k3_waybel_0(X12,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12))&((r1_orders_2(X12,esk4_3(X12,X13,X14),esk5_3(X12,X13,X14))|r2_hidden(esk4_3(X12,X13,X14),X14)|X14=k3_waybel_0(X12,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12))&(r2_hidden(esk5_3(X12,X13,X14),X13)|r2_hidden(esk4_3(X12,X13,X14),X14)|X14=k3_waybel_0(X12,X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d15_waybel_0])])])])])).
cnf(c_0_91, negated_conjecture, (v2_struct_0(X1)|r2_hidden(X2,a_2_0_waybel_0(esk6_0,X3))|~r2_hidden(esk11_3(X2,X1,X4),X3)|~r2_hidden(X2,a_2_0_waybel_0(X1,X4))|~v12_waybel_0(a_2_0_waybel_0(esk6_0,X3),X1)|~m1_subset_1(a_2_0_waybel_0(esk6_0,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk11_3(X2,X1,X4),u1_struct_0(esk6_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_75, c_0_30])).
cnf(c_0_92, negated_conjecture, (v12_waybel_0(a_2_0_waybel_0(esk6_0,X1),esk6_0)|~v12_waybel_0(X1,esk6_0)|~m1_subset_1(a_2_0_waybel_0(esk6_0,X1),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_55]), c_0_19])])).
cnf(c_0_93, negated_conjecture, (esk7_0=k5_waybel_0(esk6_0,esk8_0)|v14_waybel_0(esk7_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_94, negated_conjecture, (m1_subset_1(esk8_0,u1_struct_0(esk6_0))|v14_waybel_0(esk7_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_95, plain, (v2_struct_0(X1)|r2_hidden(X2,X3)|~r2_hidden(X2,k3_waybel_0(X1,X3))|~v12_waybel_0(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_87, c_0_31])).
cnf(c_0_96, negated_conjecture, (r2_hidden(esk10_3(X1,esk6_0,esk7_0),k3_waybel_0(esk6_0,esk7_0))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_61]), c_0_19])]), c_0_20])).
cnf(c_0_97, plain, (r1_orders_2(X1,esk4_3(X1,X2,X3),esk5_3(X1,X2,X3))|r2_hidden(esk4_3(X1,X2,X3),X3)|X3=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_98, plain, (m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(X1))|r2_hidden(esk4_3(X1,X2,X3),X3)|X3=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_99, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X1))|X3=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_100, negated_conjecture, (r2_hidden(X1,a_2_0_waybel_0(esk6_0,X2))|~r2_hidden(esk11_3(X1,esk6_0,X3),X2)|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,X3))|~v12_waybel_0(X2,esk6_0)|~m1_subset_1(a_2_0_waybel_0(esk6_0,X2),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(esk11_3(X1,esk6_0,X3),u1_struct_0(esk6_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_19])]), c_0_20])).
cnf(c_0_101, negated_conjecture, (v2_struct_0(X1)|r2_hidden(esk11_3(X2,X1,a_2_0_waybel_0(esk6_0,X3)),X3)|~r2_hidden(X2,a_2_0_waybel_0(X1,a_2_0_waybel_0(esk6_0,X3)))|~v12_waybel_0(X3,esk6_0)|~m1_subset_1(esk11_3(X2,X1,a_2_0_waybel_0(esk6_0,X3)),u1_struct_0(esk6_0))|~m1_subset_1(a_2_0_waybel_0(esk6_0,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk6_0)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_68, c_0_63])).
cnf(c_0_102, negated_conjecture, (v14_waybel_0(esk7_0,esk6_0)|r1_orders_2(esk6_0,X1,esk8_0)|~r2_hidden(X1,esk7_0)|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_93]), c_0_19])]), c_0_20]), c_0_94])).
cnf(c_0_103, negated_conjecture, (r2_hidden(esk10_3(X1,esk6_0,esk7_0),esk7_0)|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_60]), c_0_61]), c_0_19])]), c_0_20])).
cnf(c_0_104, plain, (X1=k3_waybel_0(X2,X3)|r2_hidden(esk4_3(X2,X3,X1),X1)|r2_hidden(esk4_3(X2,X3,X1),X4)|~r2_hidden(esk5_3(X2,X3,X1),X4)|~v12_waybel_0(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_97]), c_0_98]), c_0_99])).
cnf(c_0_105, plain, (r2_hidden(esk5_3(X1,X2,X3),X2)|r2_hidden(esk4_3(X1,X2,X3),X3)|X3=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_106, negated_conjecture, (r2_hidden(X1,a_2_0_waybel_0(esk6_0,X2))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,a_2_0_waybel_0(esk6_0,X2)))|~v12_waybel_0(X2,esk6_0)|~m1_subset_1(esk11_3(X1,esk6_0,a_2_0_waybel_0(esk6_0,X2)),u1_struct_0(esk6_0))|~m1_subset_1(a_2_0_waybel_0(esk6_0,X2),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_19])]), c_0_20])).
cnf(c_0_107, negated_conjecture, (r2_lattice3(esk6_0,X1,esk8_0)|v14_waybel_0(esk7_0,esk6_0)|~r2_hidden(esk9_3(esk6_0,X1,esk8_0),esk7_0)|~m1_subset_1(esk9_3(esk6_0,X1,esk8_0),u1_struct_0(esk6_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_102]), c_0_19])]), c_0_94])).
cnf(c_0_108, negated_conjecture, (r2_hidden(X1,esk7_0)|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_46]), c_0_61]), c_0_19])]), c_0_20])).
cnf(c_0_109, plain, (X1=k3_waybel_0(X2,X3)|r2_hidden(esk4_3(X2,X3,X1),X3)|r2_hidden(esk4_3(X2,X3,X1),X1)|~v12_waybel_0(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_110, negated_conjecture, (r2_hidden(X1,a_2_0_waybel_0(esk6_0,X2))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,a_2_0_waybel_0(esk6_0,X2)))|~v12_waybel_0(X2,esk6_0)|~m1_subset_1(a_2_0_waybel_0(esk6_0,X2),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_42]), c_0_19])]), c_0_20])).
cnf(c_0_111, plain, (v2_struct_0(X1)|m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X2,k3_waybel_0(X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_73, c_0_31])).
fof(c_0_112, plain, ![X42, X43, X45]:(((m1_subset_1(esk12_2(X42,X43),u1_struct_0(X42))|~v14_waybel_0(X43,X42)|(v1_xboole_0(X43)|(~v1_waybel_0(X43,X42)|(~v12_waybel_0(X43,X42)|~m1_subset_1(X43,k1_zfmisc_1(u1_struct_0(X42))))))|(v2_struct_0(X42)|(~v3_orders_2(X42)|(~v4_orders_2(X42)|~l1_orders_2(X42)))))&((r2_hidden(esk12_2(X42,X43),X43)|~v14_waybel_0(X43,X42)|(v1_xboole_0(X43)|(~v1_waybel_0(X43,X42)|(~v12_waybel_0(X43,X42)|~m1_subset_1(X43,k1_zfmisc_1(u1_struct_0(X42))))))|(v2_struct_0(X42)|(~v3_orders_2(X42)|(~v4_orders_2(X42)|~l1_orders_2(X42)))))&(r2_lattice3(X42,X43,esk12_2(X42,X43))|~v14_waybel_0(X43,X42)|(v1_xboole_0(X43)|(~v1_waybel_0(X43,X42)|(~v12_waybel_0(X43,X42)|~m1_subset_1(X43,k1_zfmisc_1(u1_struct_0(X42))))))|(v2_struct_0(X42)|(~v3_orders_2(X42)|(~v4_orders_2(X42)|~l1_orders_2(X42)))))))&(~m1_subset_1(X45,u1_struct_0(X42))|(~r2_hidden(X45,X43)|~r2_lattice3(X42,X43,X45))|v14_waybel_0(X43,X42)|(v1_xboole_0(X43)|(~v1_waybel_0(X43,X42)|(~v12_waybel_0(X43,X42)|~m1_subset_1(X43,k1_zfmisc_1(u1_struct_0(X42))))))|(v2_struct_0(X42)|(~v3_orders_2(X42)|(~v4_orders_2(X42)|~l1_orders_2(X42)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d21_waybel_0])])])])])])).
cnf(c_0_113, negated_conjecture, (r2_lattice3(esk6_0,esk7_0,esk8_0)|v14_waybel_0(esk7_0,esk6_0)|~m1_subset_1(esk9_3(esk6_0,esk7_0,esk8_0),u1_struct_0(esk6_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_33]), c_0_19])]), c_0_94])).
cnf(c_0_114, negated_conjecture, (r2_hidden(X1,k5_waybel_0(esk6_0,X1))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_25]), c_0_19])]), c_0_20])).
cnf(c_0_115, negated_conjecture, (r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))|~r2_hidden(esk11_3(X1,esk6_0,X2),esk7_0)|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,X2))|~m1_subset_1(a_2_0_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(esk11_3(X1,esk6_0,X2),u1_struct_0(esk6_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_76]), c_0_61]), c_0_19])]), c_0_20])).
cnf(c_0_116, negated_conjecture, (v2_struct_0(X1)|r2_hidden(esk11_3(X2,X1,a_2_0_waybel_0(esk6_0,esk7_0)),esk7_0)|~r2_hidden(X2,a_2_0_waybel_0(X1,a_2_0_waybel_0(esk6_0,esk7_0)))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(a_2_0_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_108, c_0_63])).
cnf(c_0_117, plain, (X3=k3_waybel_0(X1,X2)|~r2_hidden(esk4_3(X1,X2,X3),X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,esk4_3(X1,X2,X3),X4)|~r2_hidden(X4,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_118, negated_conjecture, (X1=k3_waybel_0(esk6_0,esk7_0)|r2_hidden(esk4_3(esk6_0,esk7_0,X1),esk7_0)|r2_hidden(esk4_3(esk6_0,esk7_0,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_61]), c_0_60]), c_0_19])])).
cnf(c_0_119, negated_conjecture, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_waybel_0(esk6_0,X2))|~v12_waybel_0(X2,esk6_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_31]), c_0_19])]), c_0_20])).
cnf(c_0_120, negated_conjecture, (r2_hidden(esk10_3(X1,esk6_0,a_2_0_waybel_0(esk6_0,X2)),a_2_0_waybel_0(esk6_0,X2))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,a_2_0_waybel_0(esk6_0,X2)))|~v12_waybel_0(X2,esk6_0)|~m1_subset_1(a_2_0_waybel_0(esk6_0,X2),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_89]), c_0_19])]), c_0_20])).
cnf(c_0_121, negated_conjecture, (m1_subset_1(esk10_3(X1,esk6_0,esk7_0),u1_struct_0(esk6_0))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_96]), c_0_61]), c_0_19])]), c_0_20])).
cnf(c_0_122, plain, (v14_waybel_0(X3,X2)|v1_xboole_0(X3)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r2_lattice3(X2,X3,X1)|~v1_waybel_0(X3,X2)|~v12_waybel_0(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v3_orders_2(X2)|~v4_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_123, negated_conjecture, (r2_lattice3(esk6_0,esk7_0,esk8_0)|v14_waybel_0(esk7_0,esk6_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_38]), c_0_19])]), c_0_94])).
cnf(c_0_124, negated_conjecture, (v1_waybel_0(esk7_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_125, negated_conjecture, (~v1_xboole_0(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_126, negated_conjecture, (v14_waybel_0(esk7_0,esk6_0)|r2_hidden(esk8_0,esk7_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_93]), c_0_94])).
cnf(c_0_127, negated_conjecture, (r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,a_2_0_waybel_0(esk6_0,esk7_0)))|~m1_subset_1(esk11_3(X1,esk6_0,a_2_0_waybel_0(esk6_0,esk7_0)),u1_struct_0(esk6_0))|~m1_subset_1(a_2_0_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_19])]), c_0_20])).
cnf(c_0_128, negated_conjecture, (X1=k3_waybel_0(esk6_0,X2)|~r2_hidden(esk4_3(esk6_0,X2,X1),X1)|~r2_hidden(esk4_3(esk6_0,X2,X1),X2)|~m1_subset_1(esk4_3(esk6_0,X2,X1),u1_struct_0(esk6_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_25]), c_0_19])])).
cnf(c_0_129, negated_conjecture, (k3_waybel_0(esk6_0,esk7_0)=esk7_0|r2_hidden(esk4_3(esk6_0,esk7_0,esk7_0),esk7_0)), inference(spm,[status(thm)],[c_0_118, c_0_61])).
cnf(c_0_130, negated_conjecture, (r2_hidden(X1,a_2_0_waybel_0(esk6_0,X2))|~r2_hidden(esk11_3(X1,esk6_0,X3),X2)|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,X3))|~m1_subset_1(esk11_3(X1,esk6_0,X3),u1_struct_0(esk6_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_57]), c_0_19])]), c_0_20])).
cnf(c_0_131, negated_conjecture, (v2_struct_0(X1)|r2_hidden(esk11_3(X2,X1,k3_waybel_0(esk6_0,X3)),X3)|~r2_hidden(X2,a_2_0_waybel_0(X1,k3_waybel_0(esk6_0,X3)))|~v12_waybel_0(X3,esk6_0)|~m1_subset_1(esk11_3(X2,X1,k3_waybel_0(esk6_0,X3)),u1_struct_0(esk6_0))|~m1_subset_1(k3_waybel_0(esk6_0,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk6_0)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_119, c_0_63])).
cnf(c_0_132, negated_conjecture, (m1_subset_1(esk10_3(X1,esk6_0,a_2_0_waybel_0(esk6_0,X2)),u1_struct_0(esk6_0))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,a_2_0_waybel_0(esk6_0,X2)))|~v12_waybel_0(X2,esk6_0)|~m1_subset_1(a_2_0_waybel_0(esk6_0,X2),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_120]), c_0_19])]), c_0_20])).
cnf(c_0_133, negated_conjecture, (r1_orders_2(esk6_0,esk10_3(X1,esk6_0,esk7_0),X2)|~r2_lattice3(esk6_0,esk7_0,X2)|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X2,u1_struct_0(esk6_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_96]), c_0_61])]), c_0_121])).
cnf(c_0_134, plain, (r2_lattice3(X1,X2,esk12_2(X1,X2))|v1_xboole_0(X2)|v2_struct_0(X1)|~v14_waybel_0(X2,X1)|~v1_waybel_0(X2,X1)|~v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_135, negated_conjecture, (v14_waybel_0(esk7_0,esk6_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_124]), c_0_23]), c_0_18]), c_0_60]), c_0_61]), c_0_19])]), c_0_125]), c_0_20]), c_0_94]), c_0_126])).
cnf(c_0_136, plain, (v2_struct_0(X1)|r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~r2_hidden(esk11_3(X2,X1,X4),X3)|~r2_hidden(X2,a_2_0_waybel_0(X1,X4))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_72]), c_0_73]), c_0_42])).
cnf(c_0_137, negated_conjecture, (r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,a_2_0_waybel_0(esk6_0,esk7_0)))|~m1_subset_1(a_2_0_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_42]), c_0_19])]), c_0_20])).
cnf(c_0_138, negated_conjecture, (k3_waybel_0(esk6_0,esk7_0)=esk7_0|~m1_subset_1(esk4_3(esk6_0,esk7_0,esk7_0),u1_struct_0(esk6_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129]), c_0_61])]), c_0_129])).
cnf(c_0_139, negated_conjecture, (r2_hidden(X1,a_2_0_waybel_0(esk6_0,X2))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,k3_waybel_0(esk6_0,X2)))|~v12_waybel_0(X2,esk6_0)|~m1_subset_1(esk11_3(X1,esk6_0,k3_waybel_0(esk6_0,X2)),u1_struct_0(esk6_0))|~m1_subset_1(k3_waybel_0(esk6_0,X2),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_131]), c_0_19])]), c_0_20])).
cnf(c_0_140, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk6_0))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,a_2_0_waybel_0(esk6_0,X2)))|~v12_waybel_0(X2,esk6_0)|~m1_subset_1(a_2_0_waybel_0(esk6_0,X2),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_46]), c_0_19])]), c_0_20])).
cnf(c_0_141, negated_conjecture, (r1_orders_2(esk6_0,esk10_3(X1,esk6_0,esk7_0),esk12_2(esk6_0,esk7_0))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(esk12_2(esk6_0,esk7_0),u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134]), c_0_135]), c_0_124]), c_0_23]), c_0_18]), c_0_60]), c_0_61]), c_0_19])]), c_0_125]), c_0_20])).
cnf(c_0_142, negated_conjecture, (v2_struct_0(X1)|r2_hidden(X2,a_2_0_waybel_0(X1,esk7_0))|~r2_hidden(X2,a_2_0_waybel_0(X1,a_2_0_waybel_0(esk6_0,esk7_0)))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(a_2_0_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk7_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_136, c_0_116])).
cnf(c_0_143, negated_conjecture, (r2_hidden(X1,k3_waybel_0(esk6_0,esk7_0))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,k3_waybel_0(esk6_0,esk7_0)))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_31]), c_0_61]), c_0_19])]), c_0_20])).
cnf(c_0_144, negated_conjecture, (k3_waybel_0(esk6_0,esk7_0)=esk7_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_99]), c_0_61]), c_0_19])])).
cnf(c_0_145, negated_conjecture, (~v14_waybel_0(esk7_0,esk6_0)|~m1_subset_1(X1,u1_struct_0(esk6_0))|esk7_0!=k5_waybel_0(esk6_0,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_146, negated_conjecture, (r2_hidden(X1,a_2_0_waybel_0(esk6_0,X2))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,k3_waybel_0(esk6_0,X2)))|~v12_waybel_0(X2,esk6_0)|~m1_subset_1(k3_waybel_0(esk6_0,X2),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_42]), c_0_19])]), c_0_20])).
cnf(c_0_147, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk6_0))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,k3_waybel_0(esk6_0,X2)))|~v12_waybel_0(X2,esk6_0)|~m1_subset_1(k3_waybel_0(esk6_0,X2),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_31]), c_0_19])]), c_0_20])).
cnf(c_0_148, negated_conjecture, (r1_orders_2(esk6_0,X1,esk12_2(esk6_0,esk7_0))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(esk12_2(esk6_0,esk7_0),u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_46]), c_0_61]), c_0_19])]), c_0_20])).
cnf(c_0_149, negated_conjecture, (v2_struct_0(X1)|r2_hidden(X2,a_2_0_waybel_0(X1,esk7_0))|~r2_hidden(X2,a_2_0_waybel_0(X1,k3_waybel_0(esk6_0,esk7_0)))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk7_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_31]), c_0_61]), c_0_19])]), c_0_20])).
cnf(c_0_150, negated_conjecture, (r2_hidden(esk10_3(X1,esk6_0,k3_waybel_0(esk6_0,esk7_0)),k3_waybel_0(esk6_0,esk7_0))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,k3_waybel_0(esk6_0,esk7_0)))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_89]), c_0_19])]), c_0_20])).
cnf(c_0_151, plain, (r1_orders_2(X1,X2,esk3_4(X1,X3,X4,X2))|~r2_hidden(X2,X4)|~m1_subset_1(X2,u1_struct_0(X1))|X4!=k3_waybel_0(X1,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_152, plain, (m1_subset_1(esk3_4(X1,X2,X3,X4),u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_153, negated_conjecture, (X1=esk7_0|r2_hidden(esk4_3(esk6_0,esk7_0,X1),esk7_0)|r2_hidden(esk4_3(esk6_0,esk7_0,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(rw,[status(thm)],[c_0_118, c_0_144])).
cnf(c_0_154, negated_conjecture, (k5_waybel_0(esk6_0,X1)!=esk7_0|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_145, c_0_135])])).
cnf(c_0_155, negated_conjecture, (r2_hidden(X1,a_2_0_waybel_0(esk6_0,X2))|~r2_hidden(X1,k3_waybel_0(esk6_0,k3_waybel_0(esk6_0,X2)))|~v12_waybel_0(X2,esk6_0)|~m1_subset_1(k3_waybel_0(esk6_0,X2),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_31]), c_0_19])]), c_0_20])).
cnf(c_0_156, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk6_0))|~r2_hidden(X1,k3_waybel_0(esk6_0,k3_waybel_0(esk6_0,X2)))|~v12_waybel_0(X2,esk6_0)|~m1_subset_1(k3_waybel_0(esk6_0,X2),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_31]), c_0_19])]), c_0_20])).
cnf(c_0_157, negated_conjecture, (r2_lattice3(esk6_0,X1,esk12_2(esk6_0,esk7_0))|~r2_hidden(esk9_3(esk6_0,X1,esk12_2(esk6_0,esk7_0)),a_2_0_waybel_0(esk6_0,esk7_0))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(esk12_2(esk6_0,esk7_0),u1_struct_0(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_148]), c_0_19])])).
cnf(c_0_158, negated_conjecture, (v2_struct_0(X1)|r2_hidden(X2,a_2_0_waybel_0(X1,esk7_0))|~r2_hidden(X2,k3_waybel_0(X1,k3_waybel_0(esk6_0,esk7_0)))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk7_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_149, c_0_31])).
cnf(c_0_159, negated_conjecture, (m1_subset_1(esk10_3(X1,esk6_0,k3_waybel_0(esk6_0,esk7_0)),u1_struct_0(esk6_0))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,k3_waybel_0(esk6_0,esk7_0)))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_150]), c_0_61]), c_0_19])]), c_0_20])).
cnf(c_0_160, negated_conjecture, (X1=k3_waybel_0(esk6_0,X2)|~r2_hidden(esk11_3(esk4_3(esk6_0,X2,X1),esk6_0,X3),X2)|~r2_hidden(esk4_3(esk6_0,X2,X1),a_2_0_waybel_0(esk6_0,X3))|~r2_hidden(esk4_3(esk6_0,X2,X1),X1)|~m1_subset_1(esk11_3(esk4_3(esk6_0,X2,X1),esk6_0,X3),u1_struct_0(esk6_0))|~m1_subset_1(esk4_3(esk6_0,X2,X1),u1_struct_0(esk6_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_57]), c_0_19])])).
cnf(c_0_161, plain, (v2_struct_0(X1)|r2_hidden(X2,a_2_0_waybel_0(X1,X3))|X4!=k3_waybel_0(X1,X5)|~r2_hidden(esk3_4(X1,X5,X4,X2),X3)|~r2_hidden(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_151]), c_0_152])).
cnf(c_0_162, plain, (r2_hidden(esk3_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_163, plain, (r2_hidden(X3,X5)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,X3,X1)|~r2_hidden(X1,X4)|~m1_subset_1(X3,u1_struct_0(X2))|X5!=k3_waybel_0(X2,X4)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_164, negated_conjecture, (r2_hidden(esk4_3(esk6_0,esk7_0,k5_waybel_0(esk6_0,X1)),k5_waybel_0(esk6_0,X1))|r2_hidden(esk4_3(esk6_0,esk7_0,k5_waybel_0(esk6_0,X1)),esk7_0)|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_52]), c_0_19])]), c_0_20]), c_0_154])).
cnf(c_0_165, negated_conjecture, (r2_hidden(X1,a_2_0_waybel_0(esk6_0,esk7_0))|~r2_hidden(X1,k3_waybel_0(esk6_0,k3_waybel_0(esk6_0,k3_waybel_0(esk6_0,esk7_0))))|~m1_subset_1(k3_waybel_0(esk6_0,k3_waybel_0(esk6_0,esk7_0)),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_155]), c_0_61]), c_0_19])]), c_0_20]), c_0_156]), c_0_78])).
cnf(c_0_166, negated_conjecture, (r2_lattice3(esk6_0,X1,esk12_2(esk6_0,esk7_0))|~r2_hidden(esk9_3(esk6_0,X1,esk12_2(esk6_0,esk7_0)),k3_waybel_0(esk6_0,k3_waybel_0(esk6_0,esk7_0)))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(esk12_2(esk6_0,esk7_0),u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_158]), c_0_61]), c_0_19])]), c_0_20])).
cnf(c_0_167, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk6_0))|~r2_hidden(X1,a_2_0_waybel_0(esk6_0,k3_waybel_0(esk6_0,esk7_0)))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_46]), c_0_19])]), c_0_20])).
cnf(c_0_168, negated_conjecture, (X1=k3_waybel_0(esk6_0,X2)|~r2_hidden(esk4_3(esk6_0,X2,X1),a_2_0_waybel_0(esk6_0,X2))|~r2_hidden(esk4_3(esk6_0,X2,X1),X1)|~m1_subset_1(esk11_3(esk4_3(esk6_0,X2,X1),esk6_0,X2),u1_struct_0(esk6_0))|~m1_subset_1(esk4_3(esk6_0,X2,X1),u1_struct_0(esk6_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_63]), c_0_19])]), c_0_20])).
cnf(c_0_169, plain, (v2_struct_0(X1)|r2_hidden(X2,a_2_0_waybel_0(X1,X3))|X4!=k3_waybel_0(X1,X3)|~r2_hidden(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_161, c_0_162])).
cnf(c_0_170, negated_conjecture, (r2_hidden(X1,esk7_0)|esk7_0!=k3_waybel_0(esk6_0,X2)|~r1_orders_2(esk6_0,X1,X3)|~r2_hidden(X3,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))|~m1_subset_1(X3,u1_struct_0(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_61]), c_0_19])])).
cnf(c_0_171, negated_conjecture, (r1_orders_2(esk6_0,esk4_3(esk6_0,esk7_0,k5_waybel_0(esk6_0,X1)),X1)|r2_hidden(esk4_3(esk6_0,esk7_0,k5_waybel_0(esk6_0,X1)),esk7_0)|~m1_subset_1(esk4_3(esk6_0,esk7_0,k5_waybel_0(esk6_0,X1)),u1_struct_0(esk6_0))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_164]), c_0_19])]), c_0_20])).
cnf(c_0_172, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk6_0))|~r2_hidden(X1,k3_waybel_0(esk6_0,k3_waybel_0(esk6_0,k3_waybel_0(esk6_0,esk7_0))))|~m1_subset_1(k3_waybel_0(esk6_0,k3_waybel_0(esk6_0,esk7_0)),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_165]), c_0_61]), c_0_19])]), c_0_20])).
cnf(c_0_173, negated_conjecture, (r2_lattice3(esk6_0,k3_waybel_0(esk6_0,k3_waybel_0(esk6_0,esk7_0)),esk12_2(esk6_0,esk7_0))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(esk12_2(esk6_0,esk7_0),u1_struct_0(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_33]), c_0_19])])).
cnf(c_0_174, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk6_0))|~r2_hidden(X1,k3_waybel_0(esk6_0,k3_waybel_0(esk6_0,esk7_0)))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_31]), c_0_19])]), c_0_20])).
cnf(c_0_175, negated_conjecture, (X1=k3_waybel_0(esk6_0,X2)|~r2_hidden(esk4_3(esk6_0,X2,X1),a_2_0_waybel_0(esk6_0,X2))|~r2_hidden(esk4_3(esk6_0,X2,X1),X1)|~m1_subset_1(esk4_3(esk6_0,X2,X1),u1_struct_0(esk6_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_42]), c_0_19])]), c_0_20])).
cnf(c_0_176, negated_conjecture, (r2_hidden(X1,a_2_0_waybel_0(esk6_0,X2))|esk7_0!=k3_waybel_0(esk6_0,X2)|~r2_hidden(X1,esk7_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_61]), c_0_19])]), c_0_20])).
cnf(c_0_177, negated_conjecture, (r2_hidden(esk4_3(esk6_0,esk7_0,k5_waybel_0(esk6_0,X1)),esk7_0)|esk7_0!=k3_waybel_0(esk6_0,X2)|~r2_hidden(X1,X2)|~m1_subset_1(esk4_3(esk6_0,esk7_0,k5_waybel_0(esk6_0,X1)),u1_struct_0(esk6_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(spm,[status(thm)],[c_0_170, c_0_171])).
cnf(c_0_178, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk6_0))|~r2_hidden(X1,esk7_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_172, c_0_144]), c_0_144]), c_0_144]), c_0_144]), c_0_144]), c_0_61]), c_0_144]), c_0_61])])).
cnf(c_0_179, negated_conjecture, (r1_orders_2(esk6_0,X1,esk12_2(esk6_0,esk7_0))|~r2_hidden(X1,k3_waybel_0(esk6_0,k3_waybel_0(esk6_0,esk7_0)))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(esk12_2(esk6_0,esk7_0),u1_struct_0(esk6_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_173]), c_0_19])]), c_0_174])).
cnf(c_0_180, negated_conjecture, (X1=k3_waybel_0(esk6_0,X2)|esk7_0!=k3_waybel_0(esk6_0,X2)|~r2_hidden(esk4_3(esk6_0,X2,X1),esk7_0)|~r2_hidden(esk4_3(esk6_0,X2,X1),X1)|~m1_subset_1(esk4_3(esk6_0,X2,X1),u1_struct_0(esk6_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(spm,[status(thm)],[c_0_175, c_0_176])).
cnf(c_0_181, negated_conjecture, (r2_hidden(esk4_3(esk6_0,esk7_0,k5_waybel_0(esk6_0,X1)),esk7_0)|~r2_hidden(X1,esk7_0)|~m1_subset_1(esk4_3(esk6_0,esk7_0,k5_waybel_0(esk6_0,X1)),u1_struct_0(esk6_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_144]), c_0_61])]), c_0_178])).
cnf(c_0_182, negated_conjecture, (r2_hidden(X1,k5_waybel_0(esk6_0,esk12_2(esk6_0,esk7_0)))|~r2_hidden(X1,k3_waybel_0(esk6_0,k3_waybel_0(esk6_0,esk7_0)))|~m1_subset_1(k3_waybel_0(esk6_0,esk7_0),k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(esk12_2(esk6_0,esk7_0),u1_struct_0(esk6_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_179]), c_0_19])]), c_0_20]), c_0_174])).
cnf(c_0_183, negated_conjecture, (k5_waybel_0(esk6_0,X1)=esk7_0|~r2_hidden(esk4_3(esk6_0,esk7_0,k5_waybel_0(esk6_0,X1)),k5_waybel_0(esk6_0,X1))|~r2_hidden(X1,esk7_0)|~m1_subset_1(esk4_3(esk6_0,esk7_0,k5_waybel_0(esk6_0,X1)),u1_struct_0(esk6_0))|~m1_subset_1(k5_waybel_0(esk6_0,X1),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_181]), c_0_144]), c_0_144]), c_0_61])])).
cnf(c_0_184, negated_conjecture, (r2_hidden(X1,k5_waybel_0(esk6_0,esk12_2(esk6_0,esk7_0)))|~r2_hidden(X1,esk7_0)|~m1_subset_1(esk12_2(esk6_0,esk7_0),u1_struct_0(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_182, c_0_144]), c_0_144]), c_0_144]), c_0_61])])).
cnf(c_0_185, negated_conjecture, (k5_waybel_0(esk6_0,esk12_2(esk6_0,esk7_0))=esk7_0|~r2_hidden(esk4_3(esk6_0,esk7_0,k5_waybel_0(esk6_0,esk12_2(esk6_0,esk7_0))),esk7_0)|~r2_hidden(esk12_2(esk6_0,esk7_0),esk7_0)|~m1_subset_1(k5_waybel_0(esk6_0,esk12_2(esk6_0,esk7_0)),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_184]), c_0_178]), c_0_178])).
cnf(c_0_186, negated_conjecture, (k5_waybel_0(esk6_0,esk12_2(esk6_0,esk7_0))=esk7_0|~r2_hidden(esk12_2(esk6_0,esk7_0),esk7_0)|~m1_subset_1(esk4_3(esk6_0,esk7_0,k5_waybel_0(esk6_0,esk12_2(esk6_0,esk7_0))),u1_struct_0(esk6_0))|~m1_subset_1(k5_waybel_0(esk6_0,esk12_2(esk6_0,esk7_0)),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(spm,[status(thm)],[c_0_185, c_0_181])).
cnf(c_0_187, negated_conjecture, (k5_waybel_0(esk6_0,esk12_2(esk6_0,esk7_0))=esk7_0|~r2_hidden(esk12_2(esk6_0,esk7_0),esk7_0)|~m1_subset_1(k5_waybel_0(esk6_0,esk12_2(esk6_0,esk7_0)),k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_99]), c_0_144]), c_0_61]), c_0_19])])).
cnf(c_0_188, negated_conjecture, (k5_waybel_0(esk6_0,esk12_2(esk6_0,esk7_0))=esk7_0|~r2_hidden(esk12_2(esk6_0,esk7_0),esk7_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_52]), c_0_19])]), c_0_20]), c_0_178])).
cnf(c_0_189, negated_conjecture, (~r2_hidden(esk12_2(esk6_0,esk7_0),esk7_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_188]), c_0_178])).
cnf(c_0_190, plain, (r2_hidden(esk12_2(X1,X2),X2)|v1_xboole_0(X2)|v2_struct_0(X1)|~v14_waybel_0(X2,X1)|~v1_waybel_0(X2,X1)|~v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_191, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_190]), c_0_135]), c_0_124]), c_0_23]), c_0_18]), c_0_60]), c_0_61]), c_0_19])]), c_0_125]), c_0_20]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 192
# Proof object clause steps            : 169
# Proof object formula steps           : 23
# Proof object conjectures             : 115
# Proof object clause conjectures      : 112
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 45
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 119
# Proof object simplifying inferences  : 319
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 47
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 47
# Processed clauses                    : 8415
# ...of these trivial                  : 172
# ...subsumed                          : 6256
# ...remaining for further processing  : 1987
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 346
# Backward-rewritten                   : 645
# Generated clauses                    : 51031
# ...of the previous two non-trivial   : 49783
# Contextual simplify-reflections      : 305
# Paramodulations                      : 51030
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 948
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 936
# Current number of unprocessed clauses: 41001
# ...number of literals in the above   : 457965
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1038
# Clause-clause subsumption calls (NU) : 1421415
# Rec. Clause-clause subsumption calls : 40600
# Non-unit clause-clause subsumptions  : 6906
# Unit Clause-clause subsumption calls : 435
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2780244

# -------------------------------------------------
# User time                : 1.680 s
# System time              : 0.017 s
# Total time               : 1.696 s
# Maximum resident set size: 51716 pages
