# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&(v4_waybel_0(X2,X1)&m1_yellow_0(X2,X1))))=>![X3]:((v1_waybel_0(X3,X2)&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r1_yellow_0(X1,X3)=>(X3=k1_xboole_0|(r1_yellow_0(X2,X3)&k1_yellow_0(X2,X3)=k1_yellow_0(X1,X3))))))), file('data/MPTP/proofs/6af5f479f19f4628a4e034b604ae1da3.E_input', t7_waybel_0)).
fof(d4_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_yellow_0(X2,X1)=>(v4_waybel_0(X2,X1)<=>![X3]:((v1_waybel_0(X3,X2)&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r1_yellow_0(X1,X3)=>(X3=k1_xboole_0|r2_hidden(k1_yellow_0(X1,X3),u1_struct_0(X2)))))))), file('data/MPTP/proofs/6af5f479f19f4628a4e034b604ae1da3.E_input', d4_waybel_0)).
fof(t65_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>((r1_yellow_0(X1,X3)&r2_hidden(k1_yellow_0(X1,X3),u1_struct_0(X2)))=>(r1_yellow_0(X2,X3)&k1_yellow_0(X2,X3)=k1_yellow_0(X1,X3)))))), file('data/MPTP/proofs/6af5f479f19f4628a4e034b604ae1da3.E_input', t65_yellow_0)).
fof(c_0_3, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&(v4_waybel_0(X2,X1)&m1_yellow_0(X2,X1))))=>![X3]:((v1_waybel_0(X3,X2)&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r1_yellow_0(X1,X3)=>(X3=k1_xboole_0|(r1_yellow_0(X2,X3)&k1_yellow_0(X2,X3)=k1_yellow_0(X1,X3)))))))), inference(assume_negation,[status(cth)],[t7_waybel_0])).
fof(c_0_4, plain, ![X7, X8, X9]:((~v4_waybel_0(X8,X7)|(~v1_waybel_0(X9,X8)|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X8)))|(~r1_yellow_0(X7,X9)|(X9=k1_xboole_0|r2_hidden(k1_yellow_0(X7,X9),u1_struct_0(X8)))))|~m1_yellow_0(X8,X7)|(v2_struct_0(X7)|~l1_orders_2(X7)))&(((v1_waybel_0(esk4_2(X7,X8),X8)|v4_waybel_0(X8,X7)|~m1_yellow_0(X8,X7)|(v2_struct_0(X7)|~l1_orders_2(X7)))&(m1_subset_1(esk4_2(X7,X8),k1_zfmisc_1(u1_struct_0(X8)))|v4_waybel_0(X8,X7)|~m1_yellow_0(X8,X7)|(v2_struct_0(X7)|~l1_orders_2(X7))))&((r1_yellow_0(X7,esk4_2(X7,X8))|v4_waybel_0(X8,X7)|~m1_yellow_0(X8,X7)|(v2_struct_0(X7)|~l1_orders_2(X7)))&((esk4_2(X7,X8)!=k1_xboole_0|v4_waybel_0(X8,X7)|~m1_yellow_0(X8,X7)|(v2_struct_0(X7)|~l1_orders_2(X7)))&(~r2_hidden(k1_yellow_0(X7,esk4_2(X7,X8)),u1_struct_0(X8))|v4_waybel_0(X8,X7)|~m1_yellow_0(X8,X7)|(v2_struct_0(X7)|~l1_orders_2(X7))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_waybel_0])])])])])])).
fof(c_0_5, negated_conjecture, ((~v2_struct_0(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0)))&((~v2_struct_0(esk2_0)&(v4_yellow_0(esk2_0,esk1_0)&(v4_waybel_0(esk2_0,esk1_0)&m1_yellow_0(esk2_0,esk1_0))))&((v1_waybel_0(esk3_0,esk2_0)&m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0))))&(r1_yellow_0(esk1_0,esk3_0)&(esk3_0!=k1_xboole_0&(~r1_yellow_0(esk2_0,esk3_0)|k1_yellow_0(esk2_0,esk3_0)!=k1_yellow_0(esk1_0,esk3_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_3])])])])).
fof(c_0_6, plain, ![X11, X12, X13]:((r1_yellow_0(X12,X13)|(~r1_yellow_0(X11,X13)|~r2_hidden(k1_yellow_0(X11,X13),u1_struct_0(X12)))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|(v2_struct_0(X12)|(~v4_yellow_0(X12,X11)|~m1_yellow_0(X12,X11)))|(v2_struct_0(X11)|(~v4_orders_2(X11)|~l1_orders_2(X11))))&(k1_yellow_0(X12,X13)=k1_yellow_0(X11,X13)|(~r1_yellow_0(X11,X13)|~r2_hidden(k1_yellow_0(X11,X13),u1_struct_0(X12)))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|(v2_struct_0(X12)|(~v4_yellow_0(X12,X11)|~m1_yellow_0(X12,X11)))|(v2_struct_0(X11)|(~v4_orders_2(X11)|~l1_orders_2(X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t65_yellow_0])])])])])).
cnf(c_0_7, plain, (X3=k1_xboole_0|r2_hidden(k1_yellow_0(X2,X3),u1_struct_0(X1))|v2_struct_0(X2)|~v4_waybel_0(X1,X2)|~v1_waybel_0(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_yellow_0(X2,X3)|~m1_yellow_0(X1,X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, negated_conjecture, (v1_waybel_0(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11, plain, (r1_yellow_0(X1,X2)|v2_struct_0(X1)|v2_struct_0(X3)|~r1_yellow_0(X3,X2)|~r2_hidden(k1_yellow_0(X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v4_yellow_0(X1,X3)|~m1_yellow_0(X1,X3)|~v4_orders_2(X3)|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_13, negated_conjecture, (r2_hidden(k1_yellow_0(X1,esk3_0),u1_struct_0(esk2_0))|v2_struct_0(X1)|~r1_yellow_0(X1,esk3_0)|~m1_yellow_0(esk2_0,X1)|~v4_waybel_0(esk2_0,X1)|~l1_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_9])]), c_0_10])).
cnf(c_0_14, negated_conjecture, (r1_yellow_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_15, negated_conjecture, (m1_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_16, negated_conjecture, (v4_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_17, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_18, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_19, negated_conjecture, (r1_yellow_0(esk2_0,esk3_0)|v2_struct_0(X1)|~r2_hidden(k1_yellow_0(X1,esk3_0),u1_struct_0(esk2_0))|~r1_yellow_0(X1,esk3_0)|~m1_yellow_0(esk2_0,X1)|~v4_yellow_0(esk2_0,X1)|~l1_orders_2(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_8]), c_0_12])).
cnf(c_0_20, negated_conjecture, (r2_hidden(k1_yellow_0(esk1_0,esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16]), c_0_17])]), c_0_18])).
cnf(c_0_21, negated_conjecture, (v4_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_22, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_23, plain, (k1_yellow_0(X1,X2)=k1_yellow_0(X3,X2)|v2_struct_0(X1)|v2_struct_0(X3)|~r1_yellow_0(X3,X2)|~r2_hidden(k1_yellow_0(X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v4_yellow_0(X1,X3)|~m1_yellow_0(X1,X3)|~v4_orders_2(X3)|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_24, negated_conjecture, (~r1_yellow_0(esk2_0,esk3_0)|k1_yellow_0(esk2_0,esk3_0)!=k1_yellow_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_25, negated_conjecture, (r1_yellow_0(esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_15]), c_0_20]), c_0_14]), c_0_21]), c_0_17]), c_0_22])]), c_0_18])).
cnf(c_0_26, negated_conjecture, (k1_yellow_0(esk2_0,esk3_0)=k1_yellow_0(X1,esk3_0)|v2_struct_0(X1)|~r2_hidden(k1_yellow_0(X1,esk3_0),u1_struct_0(esk2_0))|~r1_yellow_0(X1,esk3_0)|~m1_yellow_0(esk2_0,X1)|~v4_yellow_0(esk2_0,X1)|~l1_orders_2(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_8]), c_0_12])).
cnf(c_0_27, negated_conjecture, (k1_yellow_0(esk2_0,esk3_0)!=k1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25])])).
cnf(c_0_28, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_14]), c_0_20]), c_0_15]), c_0_21]), c_0_17]), c_0_22])]), c_0_27]), c_0_18]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 29
# Proof object clause steps            : 22
# Proof object formula steps           : 7
# Proof object conjectures             : 22
# Proof object clause conjectures      : 19
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 27
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 20
# Processed clauses                    : 47
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 47
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 10
# ...of the previous two non-trivial   : 7
# Contextual simplify-reflections      : 0
# Paramodulations                      : 10
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 25
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 11
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 22
# Clause-clause subsumption calls (NU) : 45
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2012

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28656 pages
