# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_yellow_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>(v2_lattice3(k2_yellow_1(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>r1_tarski(k11_lattice3(k2_yellow_1(X1),X2,X3),k3_xboole_0(X2,X3)))))), file('data/MPTP/proofs/6b1a64bee32b43b68b3f3462791849e4.E_input', t7_yellow_1)).
fof(d1_yellow_1, axiom, ![X1]:k2_yellow_1(X1)=g1_orders_2(X1,k1_yellow_1(X1)), file('data/MPTP/proofs/6b1a64bee32b43b68b3f3462791849e4.E_input', d1_yellow_1)).
fof(fc5_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&(v3_orders_2(k2_yellow_1(X1))&(v4_orders_2(k2_yellow_1(X1))&v5_orders_2(k2_yellow_1(X1))))), file('data/MPTP/proofs/6b1a64bee32b43b68b3f3462791849e4.E_input', fc5_yellow_1)).
fof(dt_k2_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&l1_orders_2(k2_yellow_1(X1))), file('data/MPTP/proofs/6b1a64bee32b43b68b3f3462791849e4.E_input', dt_k2_yellow_1)).
fof(t3_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r3_orders_2(k2_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3))))), file('data/MPTP/proofs/6b1a64bee32b43b68b3f3462791849e4.E_input', t3_yellow_1)).
fof(fc6_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>(~(v2_struct_0(k2_yellow_1(X1)))&v1_orders_2(k2_yellow_1(X1)))), file('data/MPTP/proofs/6b1a64bee32b43b68b3f3462791849e4.E_input', fc6_yellow_1)).
fof(l28_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k11_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))))))), file('data/MPTP/proofs/6b1a64bee32b43b68b3f3462791849e4.E_input', l28_lattice3)).
fof(dt_k11_lattice3, axiom, ![X1, X2, X3]:((l1_orders_2(X1)&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k11_lattice3(X1,X2,X3),u1_struct_0(X1))), file('data/MPTP/proofs/6b1a64bee32b43b68b3f3462791849e4.E_input', dt_k11_lattice3)).
fof(t15_lattice3, axiom, ![X1]:((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k11_lattice3(X1,X2,X3)=k11_lattice3(X1,X3,X2)))), file('data/MPTP/proofs/6b1a64bee32b43b68b3f3462791849e4.E_input', t15_lattice3)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/6b1a64bee32b43b68b3f3462791849e4.E_input', redefinition_r3_orders_2)).
fof(t19_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X1,X3))=>r1_tarski(X1,k3_xboole_0(X2,X3))), file('data/MPTP/proofs/6b1a64bee32b43b68b3f3462791849e4.E_input', t19_xboole_1)).
fof(c_0_11, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>(v2_lattice3(k2_yellow_1(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>r1_tarski(k11_lattice3(k2_yellow_1(X1),X2,X3),k3_xboole_0(X2,X3))))))), inference(assume_negation,[status(cth)],[t7_yellow_1])).
fof(c_0_12, negated_conjecture, (~v1_xboole_0(esk1_0)&(v2_lattice3(k2_yellow_1(esk1_0))&(m1_subset_1(esk2_0,u1_struct_0(k2_yellow_1(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(k2_yellow_1(esk1_0)))&~r1_tarski(k11_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0),k3_xboole_0(esk2_0,esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])).
fof(c_0_13, plain, ![X9]:k2_yellow_1(X9)=g1_orders_2(X9,k1_yellow_1(X9)), inference(variable_rename,[status(thm)],[d1_yellow_1])).
fof(c_0_14, plain, ![X11]:(v1_orders_2(k2_yellow_1(X11))&(v3_orders_2(k2_yellow_1(X11))&(v4_orders_2(k2_yellow_1(X11))&v5_orders_2(k2_yellow_1(X11))))), inference(variable_rename,[status(thm)],[fc5_yellow_1])).
fof(c_0_15, plain, ![X10]:(v1_orders_2(k2_yellow_1(X10))&l1_orders_2(k2_yellow_1(X10))), inference(variable_rename,[status(thm)],[dt_k2_yellow_1])).
fof(c_0_16, plain, ![X12, X13, X14]:((~r3_orders_2(k2_yellow_1(X12),X13,X14)|r1_tarski(X13,X14)|~m1_subset_1(X14,u1_struct_0(k2_yellow_1(X12)))|~m1_subset_1(X13,u1_struct_0(k2_yellow_1(X12)))|v1_xboole_0(X12))&(~r1_tarski(X13,X14)|r3_orders_2(k2_yellow_1(X12),X13,X14)|~m1_subset_1(X14,u1_struct_0(k2_yellow_1(X12)))|~m1_subset_1(X13,u1_struct_0(k2_yellow_1(X12)))|v1_xboole_0(X12))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_yellow_1])])])])])).
fof(c_0_17, plain, ![X15]:((~v2_struct_0(k2_yellow_1(X15))|v1_xboole_0(X15))&(v1_orders_2(k2_yellow_1(X15))|v1_xboole_0(X15))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc6_yellow_1])])])])).
fof(c_0_18, plain, ![X16, X17, X18, X19, X20]:(((r1_orders_2(X16,X19,X17)|X19!=k11_lattice3(X16,X17,X18)|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v5_orders_2(X16)|(~v2_lattice3(X16)|~l1_orders_2(X16)))))&((r1_orders_2(X16,X19,X18)|X19!=k11_lattice3(X16,X17,X18)|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v5_orders_2(X16)|(~v2_lattice3(X16)|~l1_orders_2(X16)))))&(~m1_subset_1(X20,u1_struct_0(X16))|(~r1_orders_2(X16,X20,X17)|~r1_orders_2(X16,X20,X18)|r1_orders_2(X16,X20,X19))|X19!=k11_lattice3(X16,X17,X18)|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v5_orders_2(X16)|(~v2_lattice3(X16)|~l1_orders_2(X16)))))))&((m1_subset_1(esk4_4(X16,X17,X18,X19),u1_struct_0(X16))|~r1_orders_2(X16,X19,X18)|~r1_orders_2(X16,X19,X17)|X19=k11_lattice3(X16,X17,X18)|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v5_orders_2(X16)|(~v2_lattice3(X16)|~l1_orders_2(X16)))))&(((r1_orders_2(X16,esk4_4(X16,X17,X18,X19),X17)|~r1_orders_2(X16,X19,X18)|~r1_orders_2(X16,X19,X17)|X19=k11_lattice3(X16,X17,X18)|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v5_orders_2(X16)|(~v2_lattice3(X16)|~l1_orders_2(X16)))))&(r1_orders_2(X16,esk4_4(X16,X17,X18,X19),X18)|~r1_orders_2(X16,X19,X18)|~r1_orders_2(X16,X19,X17)|X19=k11_lattice3(X16,X17,X18)|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v5_orders_2(X16)|(~v2_lattice3(X16)|~l1_orders_2(X16))))))&(~r1_orders_2(X16,esk4_4(X16,X17,X18,X19),X19)|~r1_orders_2(X16,X19,X18)|~r1_orders_2(X16,X19,X17)|X19=k11_lattice3(X16,X17,X18)|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v5_orders_2(X16)|(~v2_lattice3(X16)|~l1_orders_2(X16)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l28_lattice3])])])])])])).
fof(c_0_19, plain, ![X22, X23, X24]:(~l1_orders_2(X22)|(~m1_subset_1(X23,u1_struct_0(X22))|~m1_subset_1(X24,u1_struct_0(X22)))|m1_subset_1(k11_lattice3(X22,X23,X24),u1_struct_0(X22))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k11_lattice3])])).
fof(c_0_20, plain, ![X25, X26, X27]:(~v5_orders_2(X25)|(~v2_lattice3(X25)|~l1_orders_2(X25))|(~m1_subset_1(X26,u1_struct_0(X25))|(~m1_subset_1(X27,u1_struct_0(X25))|k11_lattice3(X25,X26,X27)=k11_lattice3(X25,X27,X26)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_lattice3])])])).
cnf(c_0_21, negated_conjecture, (v2_lattice3(k2_yellow_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22, plain, (k2_yellow_1(X1)=g1_orders_2(X1,k1_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23, plain, (v5_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_24, plain, (l1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25, plain, (r1_tarski(X2,X3)|v1_xboole_0(X1)|~r3_orders_2(k2_yellow_1(X1),X2,X3)|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_26, plain, ![X28, X29, X30]:((~r3_orders_2(X28,X29,X30)|r1_orders_2(X28,X29,X30)|(v2_struct_0(X28)|(~v3_orders_2(X28)|~l1_orders_2(X28))|(~m1_subset_1(X29,u1_struct_0(X28))|~m1_subset_1(X30,u1_struct_0(X28)))))&(~r1_orders_2(X28,X29,X30)|r3_orders_2(X28,X29,X30)|(v2_struct_0(X28)|(~v3_orders_2(X28)|~l1_orders_2(X28))|(~m1_subset_1(X29,u1_struct_0(X28))|~m1_subset_1(X30,u1_struct_0(X28)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_27, plain, (v3_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_28, plain, (v1_xboole_0(X1)|~v2_struct_0(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_29, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|X2!=k11_lattice3(X1,X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_30, plain, (m1_subset_1(k11_lattice3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_31, plain, (k11_lattice3(X1,X2,X3)=k11_lattice3(X1,X3,X2)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_32, negated_conjecture, (v2_lattice3(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))), inference(rw,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_33, plain, (v5_orders_2(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_23, c_0_22])).
cnf(c_0_34, plain, (l1_orders_2(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_24, c_0_22])).
cnf(c_0_35, plain, (v1_xboole_0(X1)|r1_tarski(X2,X3)|~r3_orders_2(g1_orders_2(X1,k1_yellow_1(X1)),X2,X3)|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(X1,k1_yellow_1(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_22]), c_0_22]), c_0_22])).
cnf(c_0_36, plain, (r3_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_37, plain, (v3_orders_2(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_27, c_0_22])).
cnf(c_0_38, plain, (v1_xboole_0(X1)|~v2_struct_0(g1_orders_2(X1,k1_yellow_1(X1)))), inference(rw,[status(thm)],[c_0_28, c_0_22])).
cnf(c_0_39, plain, (r1_orders_2(X1,k11_lattice3(X1,X2,X3),X3)|v2_struct_0(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v2_lattice3(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_29]), c_0_30])).
cnf(c_0_40, negated_conjecture, (k11_lattice3(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),X1,X2)=k11_lattice3(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),X2,X1)|~m1_subset_1(X2,u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33]), c_0_34])])).
cnf(c_0_41, negated_conjecture, (~r1_tarski(k11_lattice3(k2_yellow_1(esk1_0),esk2_0,esk3_0),k3_xboole_0(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_42, plain, ![X31, X32, X33]:(~r1_tarski(X31,X32)|~r1_tarski(X31,X33)|r1_tarski(X31,k3_xboole_0(X32,X33))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t19_xboole_1])])).
cnf(c_0_43, plain, (r1_tarski(X1,X2)|v1_xboole_0(X3)|~r1_orders_2(g1_orders_2(X3,k1_yellow_1(X3)),X1,X2)|~m1_subset_1(X2,u1_struct_0(g1_orders_2(X3,k1_yellow_1(X3))))|~m1_subset_1(X1,u1_struct_0(g1_orders_2(X3,k1_yellow_1(X3))))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37]), c_0_34])]), c_0_38])).
cnf(c_0_44, negated_conjecture, (r1_orders_2(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),k11_lattice3(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),X1,X2),X1)|v2_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_33]), c_0_34]), c_0_32])])).
cnf(c_0_45, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(k11_lattice3(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),X1,X2),u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_40]), c_0_34])])).
cnf(c_0_47, negated_conjecture, (~r1_tarski(k11_lattice3(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,esk3_0),k3_xboole_0(esk2_0,esk3_0))), inference(rw,[status(thm)],[c_0_41, c_0_22])).
cnf(c_0_48, plain, (r1_tarski(X1,k3_xboole_0(X2,X3))|~r1_tarski(X1,X2)|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_49, negated_conjecture, (v2_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))|r1_tarski(k11_lattice3(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),X1,X2),X1)|~m1_subset_1(X1,u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_46])).
cnf(c_0_50, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k2_yellow_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_51, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(k2_yellow_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_52, negated_conjecture, (~r1_tarski(k11_lattice3(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,esk3_0),esk3_0)|~r1_tarski(k11_lattice3(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,esk3_0),esk2_0)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_53, negated_conjecture, (v2_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))|r1_tarski(k11_lattice3(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),X1,X2),X2)|~m1_subset_1(X2,u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))), inference(spm,[status(thm)],[c_0_49, c_0_40])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))), inference(rw,[status(thm)],[c_0_50, c_0_22])).
cnf(c_0_55, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0))))), inference(rw,[status(thm)],[c_0_51, c_0_22])).
cnf(c_0_56, negated_conjecture, (v2_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))|~r1_tarski(k11_lattice3(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)),esk2_0,esk3_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54]), c_0_55])])).
cnf(c_0_57, negated_conjecture, (v2_struct_0(g1_orders_2(esk1_0,k1_yellow_1(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_49]), c_0_55]), c_0_54])])).
cnf(c_0_58, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_57]), c_0_45]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 59
# Proof object clause steps            : 36
# Proof object formula steps           : 23
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 35
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 27
# Processed clauses                    : 66
# ...of these trivial                  : 2
# ...subsumed                          : 3
# ...remaining for further processing  : 61
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 4
# Generated clauses                    : 32
# ...of the previous two non-trivial   : 28
# Contextual simplify-reflections      : 6
# Paramodulations                      : 29
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 29
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 18
# Current number of unprocessed clauses: 14
# ...number of literals in the above   : 99
# Current number of archived formulas  : 0
# Current number of archived clauses   : 30
# Clause-clause subsumption calls (NU) : 438
# Rec. Clause-clause subsumption calls : 27
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4298

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28744 pages
