# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t17_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(~(r1_tsep_1(X2,X3))=>![X4]:(m1_subset_1(X4,u1_struct_0(k2_tsep_1(X1,X2,X3)))=>(?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&X5=X4)&?[X5]:(m1_subset_1(X5,u1_struct_0(X3))&X5=X4))))))), file('data/MPTP/proofs/6b438182a5784d589ca2f01448406ae6.E_input', t17_tmap_1)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/6b438182a5784d589ca2f01448406ae6.E_input', d4_xboole_0)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/6b438182a5784d589ca2f01448406ae6.E_input', t12_setfam_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/6b438182a5784d589ca2f01448406ae6.E_input', t2_subset)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/6b438182a5784d589ca2f01448406ae6.E_input', t4_subset)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/6b438182a5784d589ca2f01448406ae6.E_input', t3_subset)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/6b438182a5784d589ca2f01448406ae6.E_input', fc2_struct_0)).
fof(d5_tsep_1, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(~(r1_tsep_1(X2,X3))=>![X4]:((~(v2_struct_0(X4))&(v1_pre_topc(X4)&m1_pre_topc(X4,X1)))=>(X4=k2_tsep_1(X1,X2,X3)<=>u1_struct_0(X4)=k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)))))))), file('data/MPTP/proofs/6b438182a5784d589ca2f01448406ae6.E_input', d5_tsep_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/6b438182a5784d589ca2f01448406ae6.E_input', d3_tarski)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/6b438182a5784d589ca2f01448406ae6.E_input', dt_l1_pre_topc)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/6b438182a5784d589ca2f01448406ae6.E_input', t17_xboole_1)).
fof(dt_k2_tsep_1, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))&(~(v2_struct_0(X3))&m1_pre_topc(X3,X1))))=>(~(v2_struct_0(k2_tsep_1(X1,X2,X3)))&(v1_pre_topc(k2_tsep_1(X1,X2,X3))&m1_pre_topc(k2_tsep_1(X1,X2,X3),X1)))), file('data/MPTP/proofs/6b438182a5784d589ca2f01448406ae6.E_input', dt_k2_tsep_1)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/6b438182a5784d589ca2f01448406ae6.E_input', dt_m1_pre_topc)).
fof(c_0_13, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(~(r1_tsep_1(X2,X3))=>![X4]:(m1_subset_1(X4,u1_struct_0(k2_tsep_1(X1,X2,X3)))=>(?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&X5=X4)&?[X5]:(m1_subset_1(X5,u1_struct_0(X3))&X5=X4)))))))), inference(assume_negation,[status(cth)],[t17_tmap_1])).
fof(c_0_14, plain, ![X38, X39, X40, X41, X42, X43, X44, X45]:((((r2_hidden(X41,X38)|~r2_hidden(X41,X40)|X40!=k3_xboole_0(X38,X39))&(r2_hidden(X41,X39)|~r2_hidden(X41,X40)|X40!=k3_xboole_0(X38,X39)))&(~r2_hidden(X42,X38)|~r2_hidden(X42,X39)|r2_hidden(X42,X40)|X40!=k3_xboole_0(X38,X39)))&((~r2_hidden(esk6_3(X43,X44,X45),X45)|(~r2_hidden(esk6_3(X43,X44,X45),X43)|~r2_hidden(esk6_3(X43,X44,X45),X44))|X45=k3_xboole_0(X43,X44))&((r2_hidden(esk6_3(X43,X44,X45),X43)|r2_hidden(esk6_3(X43,X44,X45),X45)|X45=k3_xboole_0(X43,X44))&(r2_hidden(esk6_3(X43,X44,X45),X44)|r2_hidden(esk6_3(X43,X44,X45),X45)|X45=k3_xboole_0(X43,X44))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_15, plain, ![X32, X33]:k1_setfam_1(k2_tarski(X32,X33))=k3_xboole_0(X32,X33), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_16, plain, ![X27, X28]:(~m1_subset_1(X27,X28)|(v1_xboole_0(X28)|r2_hidden(X27,X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_17, negated_conjecture, ![X10, X11]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk1_0))&(~r1_tsep_1(esk2_0,esk3_0)&(m1_subset_1(esk4_0,u1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0)))&(~m1_subset_1(X10,u1_struct_0(esk2_0))|X10!=esk4_0|(~m1_subset_1(X11,u1_struct_0(esk3_0))|X11!=esk4_0))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_13])])])])])).
fof(c_0_18, plain, ![X17, X18, X19]:(~r2_hidden(X17,X18)|~m1_subset_1(X18,k1_zfmisc_1(X19))|m1_subset_1(X17,X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_19, plain, ![X15, X16]:((~m1_subset_1(X15,k1_zfmisc_1(X16))|r1_tarski(X15,X16))&(~r1_tarski(X15,X16)|m1_subset_1(X15,k1_zfmisc_1(X16)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_20, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X4!=k3_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r2_hidden(esk6_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_23, plain, ![X20]:(v2_struct_0(X20)|~l1_struct_0(X20)|~v1_xboole_0(u1_struct_0(X20))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_24, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_26, plain, ![X34, X35, X36, X37]:((X37!=k2_tsep_1(X34,X35,X36)|u1_struct_0(X37)=k3_xboole_0(u1_struct_0(X35),u1_struct_0(X36))|(v2_struct_0(X37)|(~v1_pre_topc(X37)|~m1_pre_topc(X37,X34)))|r1_tsep_1(X35,X36)|(v2_struct_0(X36)|~m1_pre_topc(X36,X34))|(v2_struct_0(X35)|~m1_pre_topc(X35,X34))|(v2_struct_0(X34)|~l1_pre_topc(X34)))&(u1_struct_0(X37)!=k3_xboole_0(u1_struct_0(X35),u1_struct_0(X36))|X37=k2_tsep_1(X34,X35,X36)|(v2_struct_0(X37)|(~v1_pre_topc(X37)|~m1_pre_topc(X37,X34)))|r1_tsep_1(X35,X36)|(v2_struct_0(X36)|~m1_pre_topc(X36,X34))|(v2_struct_0(X35)|~m1_pre_topc(X35,X34))|(v2_struct_0(X34)|~l1_pre_topc(X34)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_tsep_1])])])])])).
cnf(c_0_27, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (r2_hidden(X1,X4)|X4!=k1_setfam_1(k2_tarski(X2,X3))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_20, c_0_21])).
fof(c_0_30, plain, ![X21, X22, X23, X24, X25]:((~r1_tarski(X21,X22)|(~r2_hidden(X23,X21)|r2_hidden(X23,X22)))&((r2_hidden(esk5_2(X24,X25),X24)|r1_tarski(X24,X25))&(~r2_hidden(esk5_2(X24,X25),X25)|r1_tarski(X24,X25)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_31, plain, (X3=k3_xboole_0(X1,X2)|~r2_hidden(esk6_3(X1,X2,X3),X3)|~r2_hidden(esk6_3(X1,X2,X3),X1)|~r2_hidden(esk6_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_32, plain, (X3=k1_setfam_1(k2_tarski(X1,X2))|r2_hidden(esk6_3(X1,X2,X3),X3)|r2_hidden(esk6_3(X1,X2,X3),X2)), inference(rw,[status(thm)],[c_0_22, c_0_21])).
cnf(c_0_33, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_34, negated_conjecture, (v1_xboole_0(u1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0)))|r2_hidden(esk4_0,u1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
fof(c_0_35, plain, ![X29]:(~l1_pre_topc(X29)|l1_struct_0(X29)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_36, plain, ![X30, X31]:r1_tarski(k3_xboole_0(X30,X31),X30), inference(variable_rename,[status(thm)],[t17_xboole_1])).
cnf(c_0_37, plain, (u1_struct_0(X1)=k3_xboole_0(u1_struct_0(X3),u1_struct_0(X4))|v2_struct_0(X1)|r1_tsep_1(X3,X4)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|X1!=k2_tsep_1(X2,X3,X4)|~v1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_38, plain, ![X12, X13, X14]:((~v2_struct_0(k2_tsep_1(X12,X13,X14))|(v2_struct_0(X12)|~l1_pre_topc(X12)|(v2_struct_0(X13)|~m1_pre_topc(X13,X12)|(v2_struct_0(X14)|~m1_pre_topc(X14,X12)))))&((v1_pre_topc(k2_tsep_1(X12,X13,X14))|(v2_struct_0(X12)|~l1_pre_topc(X12)|(v2_struct_0(X13)|~m1_pre_topc(X13,X12)|(v2_struct_0(X14)|~m1_pre_topc(X14,X12)))))&(m1_pre_topc(k2_tsep_1(X12,X13,X14),X12)|(v2_struct_0(X12)|~l1_pre_topc(X12)|(v2_struct_0(X13)|~m1_pre_topc(X13,X12)|(v2_struct_0(X14)|~m1_pre_topc(X14,X12))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_tsep_1])])])])).
cnf(c_0_39, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_40, plain, (r2_hidden(X1,k1_setfam_1(k2_tarski(X2,X3)))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_29])).
cnf(c_0_41, plain, (r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_42, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_43, plain, (X3=k1_setfam_1(k2_tarski(X1,X2))|~r2_hidden(esk6_3(X1,X2,X3),X3)|~r2_hidden(esk6_3(X1,X2,X3),X2)|~r2_hidden(esk6_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_31, c_0_21])).
cnf(c_0_44, plain, (k1_setfam_1(k2_tarski(X1,X2))=X2|r2_hidden(esk6_3(X1,X2,X2),X2)), inference(ef,[status(thm)],[c_0_32])).
cnf(c_0_45, negated_conjecture, (r2_hidden(esk4_0,u1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0)))|v2_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_46, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_47, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_48, plain, (u1_struct_0(X1)=k1_setfam_1(k2_tarski(u1_struct_0(X3),u1_struct_0(X4)))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|r1_tsep_1(X3,X4)|X1!=k2_tsep_1(X2,X3,X4)|~l1_pre_topc(X2)|~v1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)), inference(rw,[status(thm)],[c_0_37, c_0_21])).
cnf(c_0_49, plain, (m1_pre_topc(k2_tsep_1(X1,X2,X3),X1)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, plain, (v1_pre_topc(k2_tsep_1(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_51, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~v2_struct_0(k2_tsep_1(X1,X2,X3))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_52, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X3)|~r2_hidden(X1,X4)|~r1_tarski(k1_setfam_1(k2_tarski(X4,X3)),X2)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_53, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_54, plain, (k1_setfam_1(k2_tarski(X1,X2))=X2|~r2_hidden(esk6_3(X1,X2,X2),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_44])).
cnf(c_0_55, negated_conjecture, (r2_hidden(esk4_0,u1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0)))|v2_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_pre_topc(k2_tsep_1(esk1_0,esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_56, plain, (r1_tarski(k1_setfam_1(k2_tarski(X1,X2)),X1)), inference(rw,[status(thm)],[c_0_47, c_0_21])).
cnf(c_0_57, plain, (u1_struct_0(k2_tsep_1(X1,X2,X3))=k1_setfam_1(k2_tarski(u1_struct_0(X2),u1_struct_0(X3)))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_48]), c_0_49]), c_0_50]), c_0_51])).
fof(c_0_58, plain, ![X47, X48]:(~l1_pre_topc(X47)|(~m1_pre_topc(X48,X47)|l1_pre_topc(X48))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_59, negated_conjecture, (~m1_subset_1(X1,u1_struct_0(esk2_0))|X1!=esk4_0|~m1_subset_1(X2,u1_struct_0(esk3_0))|X2!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_60, plain, (m1_subset_1(X1,k1_setfam_1(k2_tarski(X2,X3)))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_61, plain, (k1_setfam_1(k2_tarski(X1,X1))=X1), inference(spm,[status(thm)],[c_0_54, c_0_44])).
cnf(c_0_62, negated_conjecture, (m1_subset_1(esk4_0,X1)|v2_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0))|~r1_tarski(u1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0)),X1)|~l1_pre_topc(k2_tsep_1(esk1_0,esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_39, c_0_55])).
cnf(c_0_63, plain, (r1_tarski(u1_struct_0(k2_tsep_1(X1,X2,X3)),u1_struct_0(X2))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_64, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_65, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_66, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_67, negated_conjecture, (~r1_tsep_1(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_68, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_69, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_70, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_71, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_72, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_73, negated_conjecture, (~m1_subset_1(esk4_0,u1_struct_0(esk3_0))|~m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_59])])).
cnf(c_0_74, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_75, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk2_0))|v2_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_pre_topc(k2_tsep_1(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64]), c_0_65]), c_0_66])]), c_0_67]), c_0_68]), c_0_69]), c_0_70])).
cnf(c_0_76, plain, (l1_pre_topc(k2_tsep_1(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_71, c_0_49])).
cnf(c_0_77, plain, (r2_hidden(X1,X2)|X3!=k1_setfam_1(k2_tarski(X4,X2))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_72, c_0_21])).
cnf(c_0_78, negated_conjecture, (~r2_hidden(esk4_0,u1_struct_0(esk3_0))|~m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_79, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk2_0))|v2_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_64]), c_0_65]), c_0_66])]), c_0_70]), c_0_69]), c_0_68])).
cnf(c_0_80, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_setfam_1(k2_tarski(X3,X2)))), inference(er,[status(thm)],[c_0_77])).
cnf(c_0_81, negated_conjecture, (v2_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0))|~r2_hidden(esk4_0,u1_struct_0(esk3_0))), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_82, plain, (r2_hidden(X1,u1_struct_0(X2))|r1_tsep_1(X3,X2)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~r2_hidden(X1,u1_struct_0(k2_tsep_1(X4,X3,X2)))|~m1_pre_topc(X2,X4)|~m1_pre_topc(X3,X4)|~l1_pre_topc(X4)), inference(spm,[status(thm)],[c_0_80, c_0_57])).
cnf(c_0_83, negated_conjecture, (~r2_hidden(esk4_0,u1_struct_0(esk3_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_81]), c_0_64]), c_0_65]), c_0_66])]), c_0_68]), c_0_69]), c_0_70])).
cnf(c_0_84, negated_conjecture, (v2_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_pre_topc(k2_tsep_1(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_55]), c_0_64]), c_0_65]), c_0_66])]), c_0_83]), c_0_67]), c_0_68]), c_0_69]), c_0_70])).
cnf(c_0_85, negated_conjecture, (v2_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_76]), c_0_64]), c_0_65]), c_0_66])]), c_0_70]), c_0_69]), c_0_68])).
cnf(c_0_86, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_85]), c_0_64]), c_0_65]), c_0_66])]), c_0_68]), c_0_69]), c_0_70]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 87
# Proof object clause steps            : 60
# Proof object formula steps           : 27
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 23
# Proof object simplifying inferences  : 60
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 32
# Processed clauses                    : 305
# ...of these trivial                  : 1
# ...subsumed                          : 81
# ...remaining for further processing  : 223
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 24
# Backward-rewritten                   : 32
# Generated clauses                    : 1106
# ...of the previous two non-trivial   : 1007
# Contextual simplify-reflections      : 11
# Paramodulations                      : 1081
# Factorizations                       : 20
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 130
#    Positive orientable unit clauses  : 16
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 108
# Current number of unprocessed clauses: 753
# ...number of literals in the above   : 4267
# Current number of archived formulas  : 0
# Current number of archived clauses   : 89
# Clause-clause subsumption calls (NU) : 14297
# Rec. Clause-clause subsumption calls : 2607
# Non-unit clause-clause subsumptions  : 111
# Unit Clause-clause subsumption calls : 116
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 16
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 29921

# -------------------------------------------------
# User time                : 0.025 s
# System time              : 0.006 s
# Total time               : 0.032 s
# Maximum resident set size: 28644 pages
