# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t67_zfmisc_1, conjecture, ![X1, X2]:(k4_xboole_0(k1_tarski(X1),X2)=k1_tarski(X1)<=>~(r2_hidden(X1,X2))), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', t67_zfmisc_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', t3_xboole_0)).
fof(l97_xboole_1, axiom, ![X1, X2]:r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2)), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', l97_xboole_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', t70_enumset1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', t100_xboole_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', t75_enumset1)).
fof(commutativity_k5_xboole_0, axiom, ![X1, X2]:k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', commutativity_k5_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', commutativity_k3_xboole_0)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', d4_xboole_0)).
fof(t16_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3)), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', t16_xboole_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', idempotence_k3_xboole_0)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', d7_xboole_0)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', t7_xboole_0)).
fof(d1_enumset1, axiom, ![X1, X2, X3, X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3)))))), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', d1_enumset1)).
fof(t125_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', t125_enumset1)).
fof(t112_xboole_1, axiom, ![X1, X2, X3]:k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k3_xboole_0(k5_xboole_0(X1,X3),X2), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', t112_xboole_1)).
fof(t60_zfmisc_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,X2)=>((r2_hidden(X3,X2)&~(X1=X3))|k3_xboole_0(k2_tarski(X1,X3),X2)=k1_tarski(X1))), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', t60_zfmisc_1)).
fof(t1_xboole_0, axiom, ![X1, X2, X3]:(r2_hidden(X1,k5_xboole_0(X2,X3))<=>~((r2_hidden(X1,X2)<=>r2_hidden(X1,X3)))), file('data/MPTP/proofs/6b695cf9867f419abe064a86785da03d.E_input', t1_xboole_0)).
fof(c_0_23, negated_conjecture, ~(![X1, X2]:(k4_xboole_0(k1_tarski(X1),X2)=k1_tarski(X1)<=>~(r2_hidden(X1,X2)))), inference(assume_negation,[status(cth)],[t67_zfmisc_1])).
fof(c_0_24, plain, ![X10, X11, X13, X14, X15]:(((r2_hidden(esk3_2(X10,X11),X10)|r1_xboole_0(X10,X11))&(r2_hidden(esk3_2(X10,X11),X11)|r1_xboole_0(X10,X11)))&(~r2_hidden(X15,X13)|~r2_hidden(X15,X14)|~r1_xboole_0(X13,X14))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_25, plain, ![X16, X17]:r1_xboole_0(k3_xboole_0(X16,X17),k5_xboole_0(X16,X17)), inference(variable_rename,[status(thm)],[l97_xboole_1])).
fof(c_0_26, negated_conjecture, ((k4_xboole_0(k1_tarski(esk1_0),esk2_0)!=k1_tarski(esk1_0)|r2_hidden(esk1_0,esk2_0))&(k4_xboole_0(k1_tarski(esk1_0),esk2_0)=k1_tarski(esk1_0)|~r2_hidden(esk1_0,esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_23])])])])).
fof(c_0_27, plain, ![X18]:k2_tarski(X18,X18)=k1_tarski(X18), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_28, plain, ![X19, X20]:k1_enumset1(X19,X19,X20)=k2_tarski(X19,X20), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_29, plain, ![X21, X22]:k4_xboole_0(X21,X22)=k5_xboole_0(X21,k3_xboole_0(X21,X22)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_30, plain, ![X23, X24, X25]:k2_enumset1(X23,X23,X24,X25)=k1_enumset1(X23,X24,X25), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_31, plain, ![X26, X27, X28, X29]:k3_enumset1(X26,X26,X27,X28,X29)=k2_enumset1(X26,X27,X28,X29), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_32, plain, ![X30, X31, X32, X33, X34]:k4_enumset1(X30,X30,X31,X32,X33,X34)=k3_enumset1(X30,X31,X32,X33,X34), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_33, plain, ![X35, X36, X37, X38, X39, X40]:k5_enumset1(X35,X35,X36,X37,X38,X39,X40)=k4_enumset1(X35,X36,X37,X38,X39,X40), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_34, plain, ![X41, X42, X43, X44, X45, X46, X47]:k6_enumset1(X41,X41,X42,X43,X44,X45,X46,X47)=k5_enumset1(X41,X42,X43,X44,X45,X46,X47), inference(variable_rename,[status(thm)],[t75_enumset1])).
cnf(c_0_35, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_36, plain, (r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_37, plain, ![X48, X49]:k5_xboole_0(X48,X49)=k5_xboole_0(X49,X48), inference(variable_rename,[status(thm)],[commutativity_k5_xboole_0])).
cnf(c_0_38, negated_conjecture, (k4_xboole_0(k1_tarski(esk1_0),esk2_0)=k1_tarski(esk1_0)|~r2_hidden(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_39, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_40, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_41, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_42, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_43, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_44, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_45, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_46, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_47, plain, ![X50, X51]:k3_xboole_0(X50,X51)=k3_xboole_0(X51,X50), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_48, plain, ![X52, X53, X54, X55, X56, X57, X58, X59]:((((r2_hidden(X55,X52)|~r2_hidden(X55,X54)|X54!=k3_xboole_0(X52,X53))&(r2_hidden(X55,X53)|~r2_hidden(X55,X54)|X54!=k3_xboole_0(X52,X53)))&(~r2_hidden(X56,X52)|~r2_hidden(X56,X53)|r2_hidden(X56,X54)|X54!=k3_xboole_0(X52,X53)))&((~r2_hidden(esk4_3(X57,X58,X59),X59)|(~r2_hidden(esk4_3(X57,X58,X59),X57)|~r2_hidden(esk4_3(X57,X58,X59),X58))|X59=k3_xboole_0(X57,X58))&((r2_hidden(esk4_3(X57,X58,X59),X57)|r2_hidden(esk4_3(X57,X58,X59),X59)|X59=k3_xboole_0(X57,X58))&(r2_hidden(esk4_3(X57,X58,X59),X58)|r2_hidden(esk4_3(X57,X58,X59),X59)|X59=k3_xboole_0(X57,X58))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_49, plain, (~r2_hidden(X1,k5_xboole_0(X2,X3))|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_50, plain, (k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_51, negated_conjecture, (k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),esk2_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)|~r2_hidden(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_39]), c_0_39]), c_0_40]), c_0_40]), c_0_41]), c_0_42]), c_0_42]), c_0_42]), c_0_43]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_46])).
cnf(c_0_52, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_53, plain, ![X61, X62, X63]:k3_xboole_0(k3_xboole_0(X61,X62),X63)=k3_xboole_0(X61,k3_xboole_0(X62,X63)), inference(variable_rename,[status(thm)],[t16_xboole_1])).
fof(c_0_54, plain, ![X64]:k3_xboole_0(X64,X64)=X64, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_55, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_56, plain, ![X65, X66]:((~r1_xboole_0(X65,X66)|k3_xboole_0(X65,X66)=k1_xboole_0)&(k3_xboole_0(X65,X66)!=k1_xboole_0|r1_xboole_0(X65,X66))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_57, plain, (~r2_hidden(X1,k5_xboole_0(X2,X3))|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_58, negated_conjecture, (k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(esk2_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)|~r2_hidden(esk1_0,esk2_0)), inference(rw,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_59, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_60, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_61, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_55])).
fof(c_0_62, plain, ![X67]:(X67=k1_xboole_0|r2_hidden(esk5_1(X67),X67)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
fof(c_0_63, plain, ![X69, X70, X71, X72, X73, X74, X75, X76, X77, X78]:(((~r2_hidden(X73,X72)|(X73=X69|(X73=X70|X73=X71))|X72!=k1_enumset1(X69,X70,X71))&((X74!=X69|r2_hidden(X74,X72)|X72!=k1_enumset1(X69,X70,X71))&((X74!=X70|r2_hidden(X74,X72)|X72!=k1_enumset1(X69,X70,X71))&(X74!=X71|r2_hidden(X74,X72)|X72!=k1_enumset1(X69,X70,X71)))))&(((esk6_4(X75,X76,X77,X78)!=X75|~r2_hidden(esk6_4(X75,X76,X77,X78),X78)|X78=k1_enumset1(X75,X76,X77))&((esk6_4(X75,X76,X77,X78)!=X76|~r2_hidden(esk6_4(X75,X76,X77,X78),X78)|X78=k1_enumset1(X75,X76,X77))&(esk6_4(X75,X76,X77,X78)!=X77|~r2_hidden(esk6_4(X75,X76,X77,X78),X78)|X78=k1_enumset1(X75,X76,X77))))&(r2_hidden(esk6_4(X75,X76,X77,X78),X78)|(esk6_4(X75,X76,X77,X78)=X75|(esk6_4(X75,X76,X77,X78)=X76|esk6_4(X75,X76,X77,X78)=X77))|X78=k1_enumset1(X75,X76,X77)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])).
cnf(c_0_64, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_65, negated_conjecture, (~r2_hidden(X1,k3_xboole_0(esk2_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))|~r2_hidden(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59]), c_0_60]), c_0_61])).
cnf(c_0_66, plain, (X1=k1_xboole_0|r2_hidden(esk5_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_67, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X5,X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_68, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|k4_xboole_0(k1_tarski(esk1_0),esk2_0)!=k1_tarski(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_69, plain, ![X80, X81, X82, X83]:k2_enumset1(X80,X81,X82,X83)=k2_enumset1(X83,X82,X81,X80), inference(variable_rename,[status(thm)],[t125_enumset1])).
cnf(c_0_70, plain, (k3_xboole_0(X1,X2)!=k1_xboole_0|~r2_hidden(X3,X2)|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_35, c_0_64])).
cnf(c_0_71, negated_conjecture, (k3_xboole_0(esk2_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))=k1_xboole_0|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_72, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k6_enumset1(X4,X4,X4,X4,X4,X4,X5,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_42]), c_0_43]), c_0_44]), c_0_45]), c_0_46])).
cnf(c_0_73, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),esk2_0))!=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_39]), c_0_39]), c_0_40]), c_0_40]), c_0_41]), c_0_42]), c_0_42]), c_0_42]), c_0_43]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_46])).
cnf(c_0_74, plain, (k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_75, negated_conjecture, (~r2_hidden(X1,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))|~r2_hidden(esk1_0,esk2_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_76, plain, (r2_hidden(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X3,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_72])])).
fof(c_0_77, plain, ![X84, X85, X86]:k5_xboole_0(k3_xboole_0(X84,X85),k3_xboole_0(X86,X85))=k3_xboole_0(k5_xboole_0(X84,X86),X85), inference(variable_rename,[status(thm)],[t112_xboole_1])).
fof(c_0_78, plain, ![X87, X88, X89]:((r2_hidden(X89,X88)|k3_xboole_0(k2_tarski(X87,X89),X88)=k1_tarski(X87)|~r2_hidden(X87,X88))&(X87!=X89|k3_xboole_0(k2_tarski(X87,X89),X88)=k1_tarski(X87)|~r2_hidden(X87,X88))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t60_zfmisc_1])])])])).
cnf(c_0_79, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(esk2_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))!=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[c_0_73, c_0_52])).
cnf(c_0_80, plain, (k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k6_enumset1(X4,X4,X4,X4,X4,X3,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_46]), c_0_46])).
cnf(c_0_81, negated_conjecture, (~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_82, plain, (k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k3_xboole_0(k5_xboole_0(X1,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_83, plain, (k3_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X1)|X1!=X2|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_84, negated_conjecture, (k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(esk2_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))!=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_81])).
cnf(c_0_85, plain, (k5_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X1,k5_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_60]), c_0_50]), c_0_52])).
cnf(c_0_86, plain, (k3_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2),X3)=k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)|X1!=X2|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_83, c_0_39]), c_0_40]), c_0_40]), c_0_42]), c_0_42]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_46]), c_0_46])).
cnf(c_0_87, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X2,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_88, negated_conjecture, (k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(esk2_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))!=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_89, plain, (k3_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),X2)=k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_86])).
fof(c_0_90, plain, ![X90, X91, X92]:(((~r2_hidden(X90,X91)|~r2_hidden(X90,X92)|~r2_hidden(X90,k5_xboole_0(X91,X92)))&(r2_hidden(X90,X91)|r2_hidden(X90,X92)|~r2_hidden(X90,k5_xboole_0(X91,X92))))&((~r2_hidden(X90,X91)|r2_hidden(X90,X92)|r2_hidden(X90,k5_xboole_0(X91,X92)))&(~r2_hidden(X90,X92)|r2_hidden(X90,X91)|r2_hidden(X90,k5_xboole_0(X91,X92))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_0])])])).
cnf(c_0_91, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k6_enumset1(X2,X2,X2,X2,X2,X2,X4,X5)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_42]), c_0_43]), c_0_44]), c_0_45]), c_0_46])).
cnf(c_0_92, negated_conjecture, (~r2_hidden(esk1_0,k5_xboole_0(esk2_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_93, plain, (r2_hidden(X1,X3)|r2_hidden(X1,k5_xboole_0(X3,X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_94, plain, (r2_hidden(X1,k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_91])])).
cnf(c_0_95, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_94])]), c_0_81]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 96
# Proof object clause steps            : 49
# Proof object formula steps           : 47
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 23
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 89
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 23
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 43
# Removed in clause preprocessing      : 8
# Initial clauses in saturation        : 35
# Processed clauses                    : 608
# ...of these trivial                  : 38
# ...subsumed                          : 414
# ...remaining for further processing  : 156
# Other redundant clauses eliminated   : 17
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 29
# Generated clauses                    : 2431
# ...of the previous two non-trivial   : 1962
# Contextual simplify-reflections      : 3
# Paramodulations                      : 2413
# Factorizations                       : 4
# Equation resolutions                 : 17
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 84
#    Positive orientable unit clauses  : 27
#    Positive unorientable unit clauses: 5
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 48
# Current number of unprocessed clauses: 1390
# ...number of literals in the above   : 3411
# Current number of archived formulas  : 0
# Current number of archived clauses   : 72
# Clause-clause subsumption calls (NU) : 968
# Rec. Clause-clause subsumption calls : 906
# Non-unit clause-clause subsumptions  : 142
# Unit Clause-clause subsumption calls : 32
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 128
# BW rewrite match successes           : 81
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 25833

# -------------------------------------------------
# User time                : 0.028 s
# System time              : 0.000 s
# Total time               : 0.028 s
# Maximum resident set size: 28724 pages
