# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t10_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r2_waybel_0(X1,X2,X3)<=>~(r1_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3)))))), file('data/MPTP/proofs/6c0e64bd4f1b4914be3ed2ded20eac0d.E_input', t10_waybel_0)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/6c0e64bd4f1b4914be3ed2ded20eac0d.E_input', redefinition_k6_subset_1)).
fof(t29_yellow_6, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>r1_waybel_0(X1,X2,u1_struct_0(X1)))), file('data/MPTP/proofs/6c0e64bd4f1b4914be3ed2ded20eac0d.E_input', t29_yellow_6)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/6c0e64bd4f1b4914be3ed2ded20eac0d.E_input', t3_boole)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/6c0e64bd4f1b4914be3ed2ded20eac0d.E_input', t5_subset)).
fof(existence_m1_subset_1, axiom, ![X1]:?[X2]:m1_subset_1(X2,X1), file('data/MPTP/proofs/6c0e64bd4f1b4914be3ed2ded20eac0d.E_input', existence_m1_subset_1)).
fof(d12_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r2_waybel_0(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X4,X5)&r2_hidden(k2_waybel_0(X1,X2,X5),X3))))))), file('data/MPTP/proofs/6c0e64bd4f1b4914be3ed2ded20eac0d.E_input', d12_waybel_0)).
fof(t8_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3, X4]:(r1_tarski(X3,X4)=>((r1_waybel_0(X1,X2,X3)=>r1_waybel_0(X1,X2,X4))&(r2_waybel_0(X1,X2,X3)=>r2_waybel_0(X1,X2,X4)))))), file('data/MPTP/proofs/6c0e64bd4f1b4914be3ed2ded20eac0d.E_input', t8_waybel_0)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/6c0e64bd4f1b4914be3ed2ded20eac0d.E_input', t2_xboole_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/6c0e64bd4f1b4914be3ed2ded20eac0d.E_input', fc1_xboole_0)).
fof(c_0_10, plain, ![X11, X12, X13]:((~r2_waybel_0(X11,X12,X13)|~r1_waybel_0(X11,X12,k6_subset_1(u1_struct_0(X11),X13))|(v2_struct_0(X12)|~l1_waybel_0(X12,X11))|(v2_struct_0(X11)|~l1_struct_0(X11)))&(r1_waybel_0(X11,X12,k6_subset_1(u1_struct_0(X11),X13))|r2_waybel_0(X11,X12,X13)|(v2_struct_0(X12)|~l1_waybel_0(X12,X11))|(v2_struct_0(X11)|~l1_struct_0(X11)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t10_waybel_0])])])])])).
fof(c_0_11, plain, ![X7, X8]:k6_subset_1(X7,X8)=k4_xboole_0(X7,X8), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
fof(c_0_12, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>r1_waybel_0(X1,X2,u1_struct_0(X1))))), inference(assume_negation,[status(cth)],[t29_yellow_6])).
cnf(c_0_13, plain, (r1_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3))|r2_waybel_0(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_15, plain, ![X6]:k4_xboole_0(X6,k1_xboole_0)=X6, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_16, plain, ![X14, X15, X16]:(~r2_hidden(X14,X15)|(~m1_subset_1(X15,k1_zfmisc_1(X16))|~v1_xboole_0(X16))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_17, plain, ![X17]:m1_subset_1(esk3_1(X17),X17), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[existence_m1_subset_1])])).
fof(c_0_18, negated_conjecture, ((~v2_struct_0(esk1_0)&l1_struct_0(esk1_0))&((~v2_struct_0(esk2_0)&(v4_orders_2(esk2_0)&(v7_waybel_0(esk2_0)&l1_waybel_0(esk2_0,esk1_0))))&~r1_waybel_0(esk1_0,esk2_0,u1_struct_0(esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])).
cnf(c_0_19, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r2_waybel_0(X1,X2,X3)|r1_waybel_0(X1,X2,k4_xboole_0(u1_struct_0(X1),X3))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)), inference(rw,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_20, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (m1_subset_1(esk3_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_23, plain, ![X19, X20, X21, X22, X24, X26]:(((m1_subset_1(esk4_4(X19,X20,X21,X22),u1_struct_0(X20))|~m1_subset_1(X22,u1_struct_0(X20))|~r2_waybel_0(X19,X20,X21)|(v2_struct_0(X20)|~l1_waybel_0(X20,X19))|(v2_struct_0(X19)|~l1_struct_0(X19)))&((r1_orders_2(X20,X22,esk4_4(X19,X20,X21,X22))|~m1_subset_1(X22,u1_struct_0(X20))|~r2_waybel_0(X19,X20,X21)|(v2_struct_0(X20)|~l1_waybel_0(X20,X19))|(v2_struct_0(X19)|~l1_struct_0(X19)))&(r2_hidden(k2_waybel_0(X19,X20,esk4_4(X19,X20,X21,X22)),X21)|~m1_subset_1(X22,u1_struct_0(X20))|~r2_waybel_0(X19,X20,X21)|(v2_struct_0(X20)|~l1_waybel_0(X20,X19))|(v2_struct_0(X19)|~l1_struct_0(X19)))))&((m1_subset_1(esk5_3(X19,X20,X24),u1_struct_0(X20))|r2_waybel_0(X19,X20,X24)|(v2_struct_0(X20)|~l1_waybel_0(X20,X19))|(v2_struct_0(X19)|~l1_struct_0(X19)))&(~m1_subset_1(X26,u1_struct_0(X20))|(~r1_orders_2(X20,esk5_3(X19,X20,X24),X26)|~r2_hidden(k2_waybel_0(X19,X20,X26),X24))|r2_waybel_0(X19,X20,X24)|(v2_struct_0(X20)|~l1_waybel_0(X20,X19))|(v2_struct_0(X19)|~l1_struct_0(X19))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d12_waybel_0])])])])])])])).
fof(c_0_24, plain, ![X27, X28, X29, X30]:((~r1_waybel_0(X27,X28,X29)|r1_waybel_0(X27,X28,X30)|~r1_tarski(X29,X30)|(v2_struct_0(X28)|~l1_waybel_0(X28,X27))|(v2_struct_0(X27)|~l1_struct_0(X27)))&(~r2_waybel_0(X27,X28,X29)|r2_waybel_0(X27,X28,X30)|~r1_tarski(X29,X30)|(v2_struct_0(X28)|~l1_waybel_0(X28,X27))|(v2_struct_0(X27)|~l1_struct_0(X27)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_waybel_0])])])])])).
cnf(c_0_25, negated_conjecture, (~r1_waybel_0(esk1_0,esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, plain, (r2_waybel_0(X1,X2,k1_xboole_0)|r1_waybel_0(X1,X2,u1_struct_0(X1))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_27, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, negated_conjecture, (l1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_29, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_30, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_31, plain, ![X31]:r1_tarski(k1_xboole_0,X31), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_32, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk3_1(k1_zfmisc_1(X1)))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_33, plain, (r2_hidden(k2_waybel_0(X1,X2,esk4_4(X1,X2,X3,X4)),X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~r2_waybel_0(X1,X2,X3)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_34, plain, (r2_waybel_0(X1,X2,X4)|v2_struct_0(X2)|v2_struct_0(X1)|~r2_waybel_0(X1,X2,X3)|~r1_tarski(X3,X4)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_35, negated_conjecture, (r2_waybel_0(esk1_0,esk2_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_28])]), c_0_29]), c_0_30])).
cnf(c_0_36, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_37, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~v1_xboole_0(X3)|~m1_subset_1(X4,u1_struct_0(X2))|~r2_waybel_0(X1,X2,esk3_1(k1_zfmisc_1(X3)))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_38, negated_conjecture, (r2_waybel_0(esk1_0,esk2_0,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_27]), c_0_28])]), c_0_30]), c_0_29])).
cnf(c_0_39, negated_conjecture, (~v1_xboole_0(X1)|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_27]), c_0_28])]), c_0_29]), c_0_30])).
cnf(c_0_40, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_41, negated_conjecture, (~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_39, c_0_22])).
cnf(c_0_42, plain, ($false), inference(sr,[status(thm)],[c_0_40, c_0_41]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 43
# Proof object clause steps            : 23
# Proof object formula steps           : 20
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 21
# Processed clauses                    : 53
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 52
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 18
# ...of the previous two non-trivial   : 14
# Contextual simplify-reflections      : 0
# Paramodulations                      : 17
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 29
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 17
# Current number of unprocessed clauses: 3
# ...number of literals in the above   : 23
# Current number of archived formulas  : 0
# Current number of archived clauses   : 24
# Clause-clause subsumption calls (NU) : 196
# Rec. Clause-clause subsumption calls : 1
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 22
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2265

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.003 s
# Total time               : 0.010 s
# Maximum resident set size: 28760 pages
