# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t41_zfmisc_1, conjecture, ![X1, X2]:~((~(X1=k1_tarski(X2))&(~(X1=k1_xboole_0)&![X3]:~((r2_hidden(X3,X1)&~(X3=X2)))))), file('data/MPTP/proofs/6c9e1ef3dae04b04b2fb3c615f2cd897.E_input', t41_zfmisc_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/6c9e1ef3dae04b04b2fb3c615f2cd897.E_input', t3_xboole_0)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/6c9e1ef3dae04b04b2fb3c615f2cd897.E_input', d5_xboole_0)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/6c9e1ef3dae04b04b2fb3c615f2cd897.E_input', t7_xboole_0)).
fof(l1_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/6c9e1ef3dae04b04b2fb3c615f2cd897.E_input', l1_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/6c9e1ef3dae04b04b2fb3c615f2cd897.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/6c9e1ef3dae04b04b2fb3c615f2cd897.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/6c9e1ef3dae04b04b2fb3c615f2cd897.E_input', t71_enumset1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/6c9e1ef3dae04b04b2fb3c615f2cd897.E_input', l32_xboole_1)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/6c9e1ef3dae04b04b2fb3c615f2cd897.E_input', t83_xboole_1)).
fof(l24_zfmisc_1, axiom, ![X1, X2]:~((r1_xboole_0(k1_tarski(X1),X2)&r2_hidden(X1,X2))), file('data/MPTP/proofs/6c9e1ef3dae04b04b2fb3c615f2cd897.E_input', l24_zfmisc_1)).
fof(t32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k4_xboole_0(X2,X1)=>X1=X2), file('data/MPTP/proofs/6c9e1ef3dae04b04b2fb3c615f2cd897.E_input', t32_xboole_1)).
fof(c_0_12, negated_conjecture, ~(![X1, X2]:~((~(X1=k1_tarski(X2))&(~(X1=k1_xboole_0)&![X3]:~((r2_hidden(X3,X1)&~(X3=X2))))))), inference(assume_negation,[status(cth)],[t41_zfmisc_1])).
fof(c_0_13, negated_conjecture, ![X7]:(esk1_0!=k1_tarski(esk2_0)&(esk1_0!=k1_xboole_0&(~r2_hidden(X7,esk1_0)|X7=esk2_0))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])])).
fof(c_0_14, plain, ![X8, X9, X11, X12, X13]:(((r2_hidden(esk3_2(X8,X9),X8)|r1_xboole_0(X8,X9))&(r2_hidden(esk3_2(X8,X9),X9)|r1_xboole_0(X8,X9)))&(~r2_hidden(X13,X11)|~r2_hidden(X13,X12)|~r1_xboole_0(X11,X12))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_15, plain, ![X14, X15, X16, X17, X18, X19, X20, X21]:((((r2_hidden(X17,X14)|~r2_hidden(X17,X16)|X16!=k4_xboole_0(X14,X15))&(~r2_hidden(X17,X15)|~r2_hidden(X17,X16)|X16!=k4_xboole_0(X14,X15)))&(~r2_hidden(X18,X14)|r2_hidden(X18,X15)|r2_hidden(X18,X16)|X16!=k4_xboole_0(X14,X15)))&((~r2_hidden(esk4_3(X19,X20,X21),X21)|(~r2_hidden(esk4_3(X19,X20,X21),X19)|r2_hidden(esk4_3(X19,X20,X21),X20))|X21=k4_xboole_0(X19,X20))&((r2_hidden(esk4_3(X19,X20,X21),X19)|r2_hidden(esk4_3(X19,X20,X21),X21)|X21=k4_xboole_0(X19,X20))&(~r2_hidden(esk4_3(X19,X20,X21),X20)|r2_hidden(esk4_3(X19,X20,X21),X21)|X21=k4_xboole_0(X19,X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_16, negated_conjecture, (X1=esk2_0|~r2_hidden(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, plain, (r2_hidden(esk3_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, plain, ![X23]:(X23=k1_xboole_0|r2_hidden(esk5_1(X23),X23)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_19, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, negated_conjecture, (esk3_2(X1,esk1_0)=esk2_0|r1_xboole_0(X1,esk1_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_22, plain, (X1=k1_xboole_0|r2_hidden(esk5_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_24, plain, ![X25, X26]:((~r1_tarski(k1_tarski(X25),X26)|r2_hidden(X25,X26))&(~r2_hidden(X25,X26)|r1_tarski(k1_tarski(X25),X26))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])).
fof(c_0_25, plain, ![X27]:k2_tarski(X27,X27)=k1_tarski(X27), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_26, plain, ![X28, X29]:k1_enumset1(X28,X28,X29)=k2_tarski(X28,X29), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_27, plain, ![X30, X31, X32]:k2_enumset1(X30,X30,X31,X32)=k1_enumset1(X30,X31,X32), inference(variable_rename,[status(thm)],[t71_enumset1])).
cnf(c_0_28, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_29, negated_conjecture, (r1_xboole_0(X1,esk1_0)|r2_hidden(esk2_0,X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_30, negated_conjecture, (esk5_1(esk1_0)=esk2_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_22]), c_0_23])).
fof(c_0_31, plain, ![X33, X34]:((k4_xboole_0(X33,X34)!=k1_xboole_0|r1_tarski(X33,X34))&(~r1_tarski(X33,X34)|k4_xboole_0(X33,X34)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_32, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36, negated_conjecture, (r1_xboole_0(k4_xboole_0(X1,X2),esk1_0)|~r2_hidden(esk2_0,X2)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_37, negated_conjecture, (r2_hidden(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_30]), c_0_23])).
cnf(c_0_38, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_39, plain, (r1_tarski(k2_enumset1(X1,X1,X1,X1),X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_33]), c_0_34]), c_0_35])).
cnf(c_0_40, negated_conjecture, (r1_xboole_0(k4_xboole_0(X1,esk1_0),esk1_0)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_41, plain, (k4_xboole_0(k2_enumset1(X1,X1,X1,X1),X2)=k1_xboole_0|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
fof(c_0_42, plain, ![X35, X36]:((~r1_xboole_0(X35,X36)|k4_xboole_0(X35,X36)=X35)&(k4_xboole_0(X35,X36)!=X35|r1_xboole_0(X35,X36))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
cnf(c_0_43, negated_conjecture, (r1_xboole_0(k1_xboole_0,esk1_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_44, plain, (k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_45, negated_conjecture, (r1_xboole_0(k1_xboole_0,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_22]), c_0_23])).
cnf(c_0_46, negated_conjecture, (k4_xboole_0(k1_xboole_0,esk1_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_47, negated_conjecture, (X1!=k1_xboole_0|~r2_hidden(X2,esk1_0)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_19, c_0_46])).
cnf(c_0_48, plain, (r2_hidden(X1,X2)|~r1_tarski(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_49, negated_conjecture, (X1!=k1_xboole_0|~r2_hidden(esk2_0,X1)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_22]), c_0_30]), c_0_23])).
cnf(c_0_50, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r2_hidden(esk3_2(X1,k4_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_28, c_0_17])).
fof(c_0_51, plain, ![X37, X38]:(~r1_xboole_0(k1_tarski(X37),X38)|~r2_hidden(X37,X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l24_zfmisc_1])])).
cnf(c_0_52, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_enumset1(X1,X1,X1,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_33]), c_0_34]), c_0_35])).
cnf(c_0_53, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_54, negated_conjecture, (r1_xboole_0(X1,esk1_0)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_49, c_0_29])).
cnf(c_0_55, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_50, c_0_20])).
cnf(c_0_56, plain, (~r1_xboole_0(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_57, plain, (r2_hidden(X1,X2)|k4_xboole_0(k2_enumset1(X1,X1,X1,X1),X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_58, negated_conjecture, (k4_xboole_0(X1,esk1_0)=X1|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_44, c_0_54])).
cnf(c_0_59, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_60, negated_conjecture, (esk3_2(esk1_0,X1)=esk2_0|r1_xboole_0(esk1_0,X1)), inference(spm,[status(thm)],[c_0_16, c_0_20])).
cnf(c_0_61, plain, (r1_xboole_0(X1,k1_xboole_0)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_55, c_0_41])).
cnf(c_0_62, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(k2_enumset1(X1,X1,X1,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_33]), c_0_34]), c_0_35])).
cnf(c_0_63, negated_conjecture, (r2_hidden(X1,esk1_0)|k2_enumset1(X1,X1,X1,X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_64, negated_conjecture, (~r2_hidden(X1,esk1_0)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_59, c_0_45])).
cnf(c_0_65, negated_conjecture, (r1_xboole_0(esk1_0,X1)|r2_hidden(esk2_0,X1)), inference(spm,[status(thm)],[c_0_17, c_0_60])).
cnf(c_0_66, plain, (X1=k1_xboole_0|r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_61, c_0_22])).
cnf(c_0_67, negated_conjecture, (k2_enumset1(X1,X1,X1,X1)!=k1_xboole_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_54]), c_0_63])).
cnf(c_0_68, plain, (r2_hidden(esk4_3(X1,X2,X3),X1)|r2_hidden(esk4_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_69, negated_conjecture, (r1_xboole_0(esk1_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_37])])).
cnf(c_0_70, plain, (r2_hidden(esk4_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk4_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_71, plain, (~r2_hidden(X1,k1_xboole_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_66]), c_0_67])).
cnf(c_0_72, negated_conjecture, (esk4_3(esk1_0,X1,X2)=esk2_0|X2=k4_xboole_0(esk1_0,X1)|r2_hidden(esk4_3(esk1_0,X1,X2),X2)), inference(spm,[status(thm)],[c_0_16, c_0_68])).
cnf(c_0_73, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r2_hidden(esk5_1(k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_28, c_0_22])).
cnf(c_0_74, negated_conjecture, (k4_xboole_0(esk1_0,k1_xboole_0)=esk1_0), inference(spm,[status(thm)],[c_0_44, c_0_69])).
cnf(c_0_75, plain, (X1=k4_xboole_0(X2,X2)|r2_hidden(esk4_3(X2,X2,X1),X1)), inference(spm,[status(thm)],[c_0_70, c_0_68])).
fof(c_0_76, plain, ![X39, X40]:(k4_xboole_0(X39,X40)!=k4_xboole_0(X40,X39)|X39=X40), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t32_xboole_1])])).
cnf(c_0_77, negated_conjecture, (esk4_3(esk1_0,X1,k1_xboole_0)=esk2_0|k4_xboole_0(esk1_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_78, negated_conjecture, (~r2_hidden(esk2_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_30]), c_0_23])).
cnf(c_0_79, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_71, c_0_75])).
cnf(c_0_80, plain, (X1=X2|k4_xboole_0(X1,X2)!=k4_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_81, negated_conjecture, (k4_xboole_0(esk1_0,X1)=k1_xboole_0|~r2_hidden(esk2_0,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_77]), c_0_78])).
cnf(c_0_82, plain, (r2_hidden(X1,k2_enumset1(X1,X1,X1,X1))), inference(spm,[status(thm)],[c_0_57, c_0_79])).
cnf(c_0_83, negated_conjecture, (esk1_0!=k1_tarski(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_84, plain, (k2_enumset1(X1,X1,X1,X1)=X2|k4_xboole_0(X2,k2_enumset1(X1,X1,X1,X1))!=k1_xboole_0|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_80, c_0_41])).
cnf(c_0_85, negated_conjecture, (k4_xboole_0(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_86, negated_conjecture, (esk1_0!=k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_83, c_0_33]), c_0_34]), c_0_35])).
cnf(c_0_87, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_37])]), c_0_86]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 88
# Proof object clause steps            : 63
# Proof object formula steps           : 25
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 39
# Proof object simplifying inferences  : 27
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 21
# Processed clauses                    : 330
# ...of these trivial                  : 21
# ...subsumed                          : 157
# ...remaining for further processing  : 152
# Other redundant clauses eliminated   : 15
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 11
# Backward-rewritten                   : 14
# Generated clauses                    : 1148
# ...of the previous two non-trivial   : 774
# Contextual simplify-reflections      : 2
# Paramodulations                      : 1123
# Factorizations                       : 6
# Equation resolutions                 : 19
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 106
#    Positive orientable unit clauses  : 19
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 81
# Current number of unprocessed clauses: 441
# ...number of literals in the above   : 1310
# Current number of archived formulas  : 0
# Current number of archived clauses   : 49
# Clause-clause subsumption calls (NU) : 1285
# Rec. Clause-clause subsumption calls : 932
# Non-unit clause-clause subsumptions  : 108
# Unit Clause-clause subsumption calls : 235
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 45
# BW rewrite match successes           : 13
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 11531

# -------------------------------------------------
# User time                : 0.014 s
# System time              : 0.003 s
# Total time               : 0.017 s
# Maximum resident set size: 28620 pages
