# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C41_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4_xboole_0, conjecture, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/6cac91736f3840a5a828747736375046.E_input', t4_xboole_0)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/6cac91736f3840a5a828747736375046.E_input', d4_xboole_0)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/6cac91736f3840a5a828747736375046.E_input', t3_xboole_0)).
fof(c_0_3, negated_conjecture, ~(![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2))))), inference(assume_negation,[status(cth)],[t4_xboole_0])).
fof(c_0_4, negated_conjecture, ![X22]:(((r2_hidden(esk5_0,k3_xboole_0(esk3_0,esk4_0))|~r1_xboole_0(esk3_0,esk4_0))&(r1_xboole_0(esk3_0,esk4_0)|~r1_xboole_0(esk3_0,esk4_0)))&((r2_hidden(esk5_0,k3_xboole_0(esk3_0,esk4_0))|~r2_hidden(X22,k3_xboole_0(esk3_0,esk4_0)))&(r1_xboole_0(esk3_0,esk4_0)|~r2_hidden(X22,k3_xboole_0(esk3_0,esk4_0))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_3])])])])])])).
fof(c_0_5, plain, ![X5, X6, X7, X8, X9, X10, X11, X12]:((((r2_hidden(X8,X5)|~r2_hidden(X8,X7)|X7!=k3_xboole_0(X5,X6))&(r2_hidden(X8,X6)|~r2_hidden(X8,X7)|X7!=k3_xboole_0(X5,X6)))&(~r2_hidden(X9,X5)|~r2_hidden(X9,X6)|r2_hidden(X9,X7)|X7!=k3_xboole_0(X5,X6)))&((~r2_hidden(esk1_3(X10,X11,X12),X12)|(~r2_hidden(esk1_3(X10,X11,X12),X10)|~r2_hidden(esk1_3(X10,X11,X12),X11))|X12=k3_xboole_0(X10,X11))&((r2_hidden(esk1_3(X10,X11,X12),X10)|r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k3_xboole_0(X10,X11))&(r2_hidden(esk1_3(X10,X11,X12),X11)|r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k3_xboole_0(X10,X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_6, negated_conjecture, (r1_xboole_0(esk3_0,esk4_0)|~r2_hidden(X1,k3_xboole_0(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, plain, (r2_hidden(esk1_3(X1,X2,X3),X2)|r2_hidden(esk1_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_8, plain, ![X14, X15, X17, X18, X19]:(((r2_hidden(esk2_2(X14,X15),X14)|r1_xboole_0(X14,X15))&(r2_hidden(esk2_2(X14,X15),X15)|r1_xboole_0(X14,X15)))&(~r2_hidden(X19,X17)|~r2_hidden(X19,X18)|~r1_xboole_0(X17,X18))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_9, negated_conjecture, (k3_xboole_0(esk3_0,esk4_0)=k3_xboole_0(X1,X2)|r1_xboole_0(esk3_0,esk4_0)|r2_hidden(esk1_3(X1,X2,k3_xboole_0(esk3_0,esk4_0)),X2)), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_10, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11, negated_conjecture, (k3_xboole_0(X1,k3_xboole_0(esk3_0,esk4_0))=k3_xboole_0(esk3_0,esk4_0)|r1_xboole_0(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_6, c_0_9])).
cnf(c_0_12, negated_conjecture, (k3_xboole_0(X1,k3_xboole_0(esk3_0,esk4_0))=k3_xboole_0(esk3_0,esk4_0)|~r2_hidden(X2,esk4_0)|~r2_hidden(X2,esk3_0)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_13, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X4!=k3_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_15, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_16, negated_conjecture, (k3_xboole_0(X1,k3_xboole_0(esk3_0,esk4_0))=k3_xboole_0(esk3_0,esk4_0)|r1_xboole_0(esk4_0,X2)|~r2_hidden(esk2_2(esk4_0,X2),esk3_0)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_17, plain, (r2_hidden(esk2_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_18, plain, (r2_hidden(X1,k3_xboole_0(X2,X3))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_15])).
cnf(c_0_20, negated_conjecture, (k3_xboole_0(X1,k3_xboole_0(esk3_0,esk4_0))=k3_xboole_0(esk3_0,esk4_0)|r1_xboole_0(esk4_0,esk3_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_21, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk2_2(X1,X2),k3_xboole_0(X3,X2))|~r2_hidden(esk2_2(X1,X2),X3)), inference(spm,[status(thm)],[c_0_18, c_0_17])).
cnf(c_0_22, negated_conjecture, (r1_xboole_0(esk4_0,esk3_0)|r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_23, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk2_2(X1,X2),k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_21, c_0_13])).
cnf(c_0_24, negated_conjecture, (r1_xboole_0(esk3_0,esk4_0)|r1_xboole_0(esk4_0,esk3_0)|r2_hidden(esk2_2(esk3_0,esk4_0),X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_25, negated_conjecture, (r2_hidden(esk5_0,k3_xboole_0(esk3_0,esk4_0))|~r1_xboole_0(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_26, negated_conjecture, (r1_xboole_0(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_6, c_0_23])).
cnf(c_0_27, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_28, negated_conjecture, (r1_xboole_0(esk4_0,esk3_0)|r1_xboole_0(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_6, c_0_24])).
cnf(c_0_29, negated_conjecture, (r2_hidden(esk5_0,k3_xboole_0(esk3_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_26])])).
cnf(c_0_30, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_31, negated_conjecture, (~r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_28]), c_0_10])).
cnf(c_0_32, negated_conjecture, (r2_hidden(esk5_0,esk3_0)), inference(spm,[status(thm)],[c_0_19, c_0_29])).
cnf(c_0_33, negated_conjecture, (r2_hidden(esk5_0,esk4_0)), inference(spm,[status(thm)],[c_0_30, c_0_29])).
cnf(c_0_34, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 35
# Proof object clause steps            : 28
# Proof object formula steps           : 7
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 5
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 13
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 12
# Processed clauses                    : 4955
# ...of these trivial                  : 73
# ...subsumed                          : 4387
# ...remaining for further processing  : 495
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 40
# Backward-rewritten                   : 112
# Generated clauses                    : 44715
# ...of the previous two non-trivial   : 44299
# Contextual simplify-reflections      : 4
# Paramodulations                      : 44584
# Factorizations                       : 104
# Equation resolutions                 : 27
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 343
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 334
# Current number of unprocessed clauses: 39225
# ...number of literals in the above   : 161798
# Current number of archived formulas  : 0
# Current number of archived clauses   : 152
# Clause-clause subsumption calls (NU) : 78312
# Rec. Clause-clause subsumption calls : 42321
# Non-unit clause-clause subsumptions  : 4422
# Unit Clause-clause subsumption calls : 709
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 60
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 901766

# -------------------------------------------------
# User time                : 0.400 s
# System time              : 0.000 s
# Total time               : 0.400 s
# Maximum resident set size: 29968 pages
