# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t53_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>(![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>r1_yellow_0(X1,X2))=>v3_lattice3(X1))), file('data/MPTP/proofs/6cb5dcd74dbf424596f9b9d3e1b0bcfe.E_input', t53_yellow_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/6cb5dcd74dbf424596f9b9d3e1b0bcfe.E_input', t3_subset)).
fof(t108_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(k3_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/6cb5dcd74dbf424596f9b9d3e1b0bcfe.E_input', t108_xboole_1)).
fof(d7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X5)=>r1_orders_2(X1,X4,X5))))=>X4=X3))))))), file('data/MPTP/proofs/6cb5dcd74dbf424596f9b9d3e1b0bcfe.E_input', d7_yellow_0)).
fof(d12_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v3_lattice3(X1)<=>![X2]:?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4))))))), file('data/MPTP/proofs/6cb5dcd74dbf424596f9b9d3e1b0bcfe.E_input', d12_lattice3)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/6cb5dcd74dbf424596f9b9d3e1b0bcfe.E_input', commutativity_k3_xboole_0)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/6cb5dcd74dbf424596f9b9d3e1b0bcfe.E_input', d10_xboole_0)).
fof(t50_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((r1_yellow_0(X1,X2)=>r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))&((r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))=>r1_yellow_0(X1,X2))&((r2_yellow_0(X1,X2)=>r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))&(r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))=>r2_yellow_0(X1,X2)))))), file('data/MPTP/proofs/6cb5dcd74dbf424596f9b9d3e1b0bcfe.E_input', t50_yellow_0)).
fof(c_0_8, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>(![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>r1_yellow_0(X1,X2))=>v3_lattice3(X1)))), inference(assume_negation,[status(cth)],[t53_yellow_0])).
fof(c_0_9, negated_conjecture, ![X7]:((~v2_struct_0(esk1_0)&l1_orders_2(esk1_0))&((~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_yellow_0(esk1_0,X7))&~v3_lattice3(esk1_0))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])).
fof(c_0_10, plain, ![X22, X23]:((~m1_subset_1(X22,k1_zfmisc_1(X23))|r1_tarski(X22,X23))&(~r1_tarski(X22,X23)|m1_subset_1(X22,k1_zfmisc_1(X23)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_11, negated_conjecture, (r1_yellow_0(esk1_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_13, plain, ![X19, X20, X21]:(~r1_tarski(X19,X20)|r1_tarski(k3_xboole_0(X19,X21),X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t108_xboole_1])])).
fof(c_0_14, plain, ![X24, X25, X27, X28, X30, X31, X34]:(((m1_subset_1(esk5_2(X24,X25),u1_struct_0(X24))|~r1_yellow_0(X24,X25)|~l1_orders_2(X24))&((r2_lattice3(X24,X25,esk5_2(X24,X25))|~r1_yellow_0(X24,X25)|~l1_orders_2(X24))&((~m1_subset_1(X27,u1_struct_0(X24))|(~r2_lattice3(X24,X25,X27)|r1_orders_2(X24,esk5_2(X24,X25),X27))|~r1_yellow_0(X24,X25)|~l1_orders_2(X24))&((m1_subset_1(esk6_3(X24,X25,X28),u1_struct_0(X24))|~r2_lattice3(X24,X25,X28)|X28=esk5_2(X24,X25)|~m1_subset_1(X28,u1_struct_0(X24))|~r1_yellow_0(X24,X25)|~l1_orders_2(X24))&((r2_lattice3(X24,X25,esk6_3(X24,X25,X28))|~r2_lattice3(X24,X25,X28)|X28=esk5_2(X24,X25)|~m1_subset_1(X28,u1_struct_0(X24))|~r1_yellow_0(X24,X25)|~l1_orders_2(X24))&(~r1_orders_2(X24,X28,esk6_3(X24,X25,X28))|~r2_lattice3(X24,X25,X28)|X28=esk5_2(X24,X25)|~m1_subset_1(X28,u1_struct_0(X24))|~r1_yellow_0(X24,X25)|~l1_orders_2(X24)))))))&(((m1_subset_1(esk8_3(X24,X30,X31),u1_struct_0(X24))|m1_subset_1(esk7_3(X24,X30,X31),u1_struct_0(X24))|~r2_lattice3(X24,X30,X31)|~m1_subset_1(X31,u1_struct_0(X24))|r1_yellow_0(X24,X30)|~l1_orders_2(X24))&(((r2_lattice3(X24,X30,esk8_3(X24,X30,X31))|m1_subset_1(esk7_3(X24,X30,X31),u1_struct_0(X24))|~r2_lattice3(X24,X30,X31)|~m1_subset_1(X31,u1_struct_0(X24))|r1_yellow_0(X24,X30)|~l1_orders_2(X24))&(~m1_subset_1(X34,u1_struct_0(X24))|(~r2_lattice3(X24,X30,X34)|r1_orders_2(X24,esk8_3(X24,X30,X31),X34))|m1_subset_1(esk7_3(X24,X30,X31),u1_struct_0(X24))|~r2_lattice3(X24,X30,X31)|~m1_subset_1(X31,u1_struct_0(X24))|r1_yellow_0(X24,X30)|~l1_orders_2(X24)))&(esk8_3(X24,X30,X31)!=X31|m1_subset_1(esk7_3(X24,X30,X31),u1_struct_0(X24))|~r2_lattice3(X24,X30,X31)|~m1_subset_1(X31,u1_struct_0(X24))|r1_yellow_0(X24,X30)|~l1_orders_2(X24))))&(((m1_subset_1(esk8_3(X24,X30,X31),u1_struct_0(X24))|r2_lattice3(X24,X30,esk7_3(X24,X30,X31))|~r2_lattice3(X24,X30,X31)|~m1_subset_1(X31,u1_struct_0(X24))|r1_yellow_0(X24,X30)|~l1_orders_2(X24))&(((r2_lattice3(X24,X30,esk8_3(X24,X30,X31))|r2_lattice3(X24,X30,esk7_3(X24,X30,X31))|~r2_lattice3(X24,X30,X31)|~m1_subset_1(X31,u1_struct_0(X24))|r1_yellow_0(X24,X30)|~l1_orders_2(X24))&(~m1_subset_1(X34,u1_struct_0(X24))|(~r2_lattice3(X24,X30,X34)|r1_orders_2(X24,esk8_3(X24,X30,X31),X34))|r2_lattice3(X24,X30,esk7_3(X24,X30,X31))|~r2_lattice3(X24,X30,X31)|~m1_subset_1(X31,u1_struct_0(X24))|r1_yellow_0(X24,X30)|~l1_orders_2(X24)))&(esk8_3(X24,X30,X31)!=X31|r2_lattice3(X24,X30,esk7_3(X24,X30,X31))|~r2_lattice3(X24,X30,X31)|~m1_subset_1(X31,u1_struct_0(X24))|r1_yellow_0(X24,X30)|~l1_orders_2(X24))))&((m1_subset_1(esk8_3(X24,X30,X31),u1_struct_0(X24))|~r1_orders_2(X24,X31,esk7_3(X24,X30,X31))|~r2_lattice3(X24,X30,X31)|~m1_subset_1(X31,u1_struct_0(X24))|r1_yellow_0(X24,X30)|~l1_orders_2(X24))&(((r2_lattice3(X24,X30,esk8_3(X24,X30,X31))|~r1_orders_2(X24,X31,esk7_3(X24,X30,X31))|~r2_lattice3(X24,X30,X31)|~m1_subset_1(X31,u1_struct_0(X24))|r1_yellow_0(X24,X30)|~l1_orders_2(X24))&(~m1_subset_1(X34,u1_struct_0(X24))|(~r2_lattice3(X24,X30,X34)|r1_orders_2(X24,esk8_3(X24,X30,X31),X34))|~r1_orders_2(X24,X31,esk7_3(X24,X30,X31))|~r2_lattice3(X24,X30,X31)|~m1_subset_1(X31,u1_struct_0(X24))|r1_yellow_0(X24,X30)|~l1_orders_2(X24)))&(esk8_3(X24,X30,X31)!=X31|~r1_orders_2(X24,X31,esk7_3(X24,X30,X31))|~r2_lattice3(X24,X30,X31)|~m1_subset_1(X31,u1_struct_0(X24))|r1_yellow_0(X24,X30)|~l1_orders_2(X24))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_yellow_0])])])])])])).
fof(c_0_15, plain, ![X8, X9, X11, X13]:(((m1_subset_1(esk2_2(X8,X9),u1_struct_0(X8))|~v3_lattice3(X8)|~l1_orders_2(X8))&((r2_lattice3(X8,X9,esk2_2(X8,X9))|~v3_lattice3(X8)|~l1_orders_2(X8))&(~m1_subset_1(X11,u1_struct_0(X8))|(~r2_lattice3(X8,X9,X11)|r1_orders_2(X8,esk2_2(X8,X9),X11))|~v3_lattice3(X8)|~l1_orders_2(X8))))&((m1_subset_1(esk4_2(X8,X13),u1_struct_0(X8))|~r2_lattice3(X8,esk3_1(X8),X13)|~m1_subset_1(X13,u1_struct_0(X8))|v3_lattice3(X8)|~l1_orders_2(X8))&((r2_lattice3(X8,esk3_1(X8),esk4_2(X8,X13))|~r2_lattice3(X8,esk3_1(X8),X13)|~m1_subset_1(X13,u1_struct_0(X8))|v3_lattice3(X8)|~l1_orders_2(X8))&(~r1_orders_2(X8,X13,esk4_2(X8,X13))|~r2_lattice3(X8,esk3_1(X8),X13)|~m1_subset_1(X13,u1_struct_0(X8))|v3_lattice3(X8)|~l1_orders_2(X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_lattice3])])])])])).
cnf(c_0_16, negated_conjecture, (r1_yellow_0(esk1_0,X1)|~r1_tarski(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_17, plain, (r1_tarski(k3_xboole_0(X1,X3),X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_18, plain, ![X17, X18]:k3_xboole_0(X17,X18)=k3_xboole_0(X18,X17), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_19, plain, ![X35, X36]:(((r1_tarski(X35,X36)|X35!=X36)&(r1_tarski(X36,X35)|X35!=X36))&(~r1_tarski(X35,X36)|~r1_tarski(X36,X35)|X35=X36)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_20, plain, (r1_orders_2(X2,esk5_2(X2,X3),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r1_yellow_0(X2,X3)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (m1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|v3_lattice3(X1)|~r2_lattice3(X1,esk3_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_22, plain, ![X15, X16]:((~r1_yellow_0(X15,X16)|r1_yellow_0(X15,k3_xboole_0(X16,u1_struct_0(X15)))|(v2_struct_0(X15)|~l1_orders_2(X15)))&((~r1_yellow_0(X15,k3_xboole_0(X16,u1_struct_0(X15)))|r1_yellow_0(X15,X16)|(v2_struct_0(X15)|~l1_orders_2(X15)))&((~r2_yellow_0(X15,X16)|r2_yellow_0(X15,k3_xboole_0(X16,u1_struct_0(X15)))|(v2_struct_0(X15)|~l1_orders_2(X15)))&(~r2_yellow_0(X15,k3_xboole_0(X16,u1_struct_0(X15)))|r2_yellow_0(X15,X16)|(v2_struct_0(X15)|~l1_orders_2(X15)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_yellow_0])])])])])).
cnf(c_0_23, negated_conjecture, (r1_yellow_0(esk1_0,k3_xboole_0(X1,X2))|~r1_tarski(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_24, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (r1_orders_2(X1,esk5_2(X1,X2),esk4_2(X1,X3))|v3_lattice3(X1)|~r2_lattice3(X1,X2,esk4_2(X1,X3))|~r2_lattice3(X1,esk3_1(X1),X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_28, plain, (r2_lattice3(X1,esk3_1(X1),esk4_2(X1,X2))|v3_lattice3(X1)|~r2_lattice3(X1,esk3_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_29, plain, (r1_yellow_0(X1,X2)|v2_struct_0(X1)|~r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, negated_conjecture, (r1_yellow_0(esk1_0,k3_xboole_0(X1,X2))|~r1_tarski(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_31, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_32, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_25])).
cnf(c_0_33, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_34, plain, (v3_lattice3(X1)|~r1_orders_2(X1,X2,esk4_2(X1,X2))|~r2_lattice3(X1,esk3_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_35, plain, (r1_orders_2(X1,esk5_2(X1,X2),esk4_2(X1,esk5_2(X1,X3)))|v3_lattice3(X1)|~r2_lattice3(X1,X2,esk4_2(X1,esk5_2(X1,X3)))|~r2_lattice3(X1,esk3_1(X1),esk5_2(X1,X3))|~r1_yellow_0(X1,X2)|~r1_yellow_0(X1,X3)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_36, plain, (r2_lattice3(X1,esk3_1(X1),esk4_2(X1,esk5_2(X1,X2)))|v3_lattice3(X1)|~r2_lattice3(X1,esk3_1(X1),esk5_2(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_28, c_0_27])).
cnf(c_0_37, negated_conjecture, (r1_yellow_0(esk1_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_32])]), c_0_33])).
cnf(c_0_38, negated_conjecture, (~v3_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_39, plain, (r2_lattice3(X1,X2,esk5_2(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_40, plain, (v3_lattice3(X1)|~r2_lattice3(X1,X2,esk4_2(X1,esk5_2(X1,X2)))|~r2_lattice3(X1,esk3_1(X1),esk5_2(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_27])).
cnf(c_0_41, negated_conjecture, (r2_lattice3(esk1_0,esk3_1(esk1_0),esk4_2(esk1_0,esk5_2(esk1_0,X1)))|~r2_lattice3(esk1_0,esk3_1(esk1_0),esk5_2(esk1_0,X1))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_31])]), c_0_38])).
cnf(c_0_42, negated_conjecture, (r2_lattice3(esk1_0,X1,esk5_2(esk1_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_37]), c_0_31])])).
cnf(c_0_43, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42]), c_0_37]), c_0_31])]), c_0_38]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 44
# Proof object clause steps            : 27
# Proof object formula steps           : 17
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 16
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 39
# Processed clauses                    : 550
# ...of these trivial                  : 1
# ...subsumed                          : 321
# ...remaining for further processing  : 228
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 6
# Generated clauses                    : 837
# ...of the previous two non-trivial   : 791
# Contextual simplify-reflections      : 4
# Paramodulations                      : 835
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 182
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 175
# Current number of unprocessed clauses: 313
# ...number of literals in the above   : 2177
# Current number of archived formulas  : 0
# Current number of archived clauses   : 44
# Clause-clause subsumption calls (NU) : 14494
# Rec. Clause-clause subsumption calls : 6546
# Non-unit clause-clause subsumptions  : 322
# Unit Clause-clause subsumption calls : 65
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 21965

# -------------------------------------------------
# User time                : 0.029 s
# System time              : 0.000 s
# Total time               : 0.029 s
# Maximum resident set size: 28692 pages
