# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t28_connsp_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>~((![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X4,X3)<=>(v3_pre_topc(X4,X1)&(v4_pre_topc(X4,X1)&r2_hidden(X2,X4)))))&X3=k1_xboole_0))))), file('data/MPTP/proofs/6ced0bd2de624a3db9a437c3249d1e49.E_input', t28_connsp_2)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/6ced0bd2de624a3db9a437c3249d1e49.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/6ced0bd2de624a3db9a437c3249d1e49.E_input', t2_xboole_1)).
fof(fc4_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>v4_pre_topc(k2_struct_0(X1),X1)), file('data/MPTP/proofs/6ced0bd2de624a3db9a437c3249d1e49.E_input', fc4_pre_topc)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/6ced0bd2de624a3db9a437c3249d1e49.E_input', d3_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/6ced0bd2de624a3db9a437c3249d1e49.E_input', dt_l1_pre_topc)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/6ced0bd2de624a3db9a437c3249d1e49.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/6ced0bd2de624a3db9a437c3249d1e49.E_input', d4_subset_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/6ced0bd2de624a3db9a437c3249d1e49.E_input', t2_subset)).
fof(d5_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>r2_hidden(X2,u1_pre_topc(X1))))), file('data/MPTP/proofs/6ced0bd2de624a3db9a437c3249d1e49.E_input', d5_pre_topc)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/6ced0bd2de624a3db9a437c3249d1e49.E_input', l13_xboole_0)).
fof(d1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1))))))))), file('data/MPTP/proofs/6ced0bd2de624a3db9a437c3249d1e49.E_input', d1_pre_topc)).
fof(rc5_tops_1, axiom, ![X1]:(l1_pre_topc(X1)=>?[X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))&v2_tops_1(X2,X1))), file('data/MPTP/proofs/6ced0bd2de624a3db9a437c3249d1e49.E_input', rc5_tops_1)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/6ced0bd2de624a3db9a437c3249d1e49.E_input', t5_subset)).
fof(t9_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k4_tarski(X3,X4))))))), file('data/MPTP/proofs/6ced0bd2de624a3db9a437c3249d1e49.E_input', t9_mcart_1)).
fof(fc14_tops_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>~(v2_tops_1(k2_struct_0(X1),X1))), file('data/MPTP/proofs/6ced0bd2de624a3db9a437c3249d1e49.E_input', fc14_tops_1)).
fof(c_0_16, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>~((![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X4,X3)<=>(v3_pre_topc(X4,X1)&(v4_pre_topc(X4,X1)&r2_hidden(X2,X4)))))&X3=k1_xboole_0)))))), inference(assume_negation,[status(cth)],[t28_connsp_2])).
fof(c_0_17, negated_conjecture, ![X8]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))&((((v3_pre_topc(X8,esk1_0)|~r2_hidden(X8,esk3_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(esk1_0))))&((v4_pre_topc(X8,esk1_0)|~r2_hidden(X8,esk3_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(esk1_0))))&(r2_hidden(esk2_0,X8)|~r2_hidden(X8,esk3_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(esk1_0))))))&(~v3_pre_topc(X8,esk1_0)|(~v4_pre_topc(X8,esk1_0)|~r2_hidden(esk2_0,X8))|r2_hidden(X8,esk3_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(esk1_0)))))&esk3_0=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_16])])])])])])).
fof(c_0_18, plain, ![X9, X10]:(~r2_hidden(X9,X10)|~r1_tarski(X10,X9)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_19, plain, ![X11]:r1_tarski(k1_xboole_0,X11), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_20, negated_conjecture, (r2_hidden(X1,esk3_0)|~v3_pre_topc(X1,esk1_0)|~v4_pre_topc(X1,esk1_0)|~r2_hidden(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21, negated_conjecture, (esk3_0=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, negated_conjecture, (r2_hidden(X1,k1_xboole_0)|~v4_pre_topc(X1,esk1_0)|~v3_pre_topc(X1,esk1_0)|~r2_hidden(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_25, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
fof(c_0_26, plain, ![X12]:(~v2_pre_topc(X12)|~l1_pre_topc(X12)|v4_pre_topc(k2_struct_0(X12),X12)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc4_pre_topc])])).
fof(c_0_27, plain, ![X13]:(~l1_struct_0(X13)|k2_struct_0(X13)=u1_struct_0(X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
fof(c_0_28, plain, ![X14]:(~l1_pre_topc(X14)|l1_struct_0(X14)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_29, plain, ![X15]:m1_subset_1(k2_subset_1(X15),k1_zfmisc_1(X15)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_30, plain, ![X16]:k2_subset_1(X16)=X16, inference(variable_rename,[status(thm)],[d4_subset_1])).
cnf(c_0_31, negated_conjecture, (~v4_pre_topc(X1,esk1_0)|~v3_pre_topc(X1,esk1_0)|~r2_hidden(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_32, plain, (v4_pre_topc(k2_struct_0(X1),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_34, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_35, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_39, plain, ![X17, X18]:(~m1_subset_1(X17,X18)|(v1_xboole_0(X18)|r2_hidden(X17,X18))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_40, negated_conjecture, (~v3_pre_topc(k2_struct_0(esk1_0),esk1_0)|~r2_hidden(esk2_0,k2_struct_0(esk1_0))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33]), c_0_34])])).
cnf(c_0_41, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_42, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_37, c_0_38])).
fof(c_0_43, plain, ![X19, X20]:((~v3_pre_topc(X20,X19)|r2_hidden(X20,u1_pre_topc(X19))|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|~l1_pre_topc(X19))&(~r2_hidden(X20,u1_pre_topc(X19))|v3_pre_topc(X20,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|~l1_pre_topc(X19))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_pre_topc])])])])).
fof(c_0_44, plain, ![X21]:(~v1_xboole_0(X21)|X21=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_45, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_47, negated_conjecture, (~v3_pre_topc(u1_struct_0(esk1_0),esk1_0)|~r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42]), c_0_33])])).
cnf(c_0_48, plain, (v3_pre_topc(X1,X2)|~r2_hidden(X1,u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_49, plain, ![X22, X23, X24, X25]:(((r2_hidden(u1_struct_0(X22),u1_pre_topc(X22))|~v2_pre_topc(X22)|~l1_pre_topc(X22))&((~m1_subset_1(X23,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X22))))|(~r1_tarski(X23,u1_pre_topc(X22))|r2_hidden(k5_setfam_1(u1_struct_0(X22),X23),u1_pre_topc(X22)))|~v2_pre_topc(X22)|~l1_pre_topc(X22))&(~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X22)))|(~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X22)))|(~r2_hidden(X24,u1_pre_topc(X22))|~r2_hidden(X25,u1_pre_topc(X22))|r2_hidden(k9_subset_1(u1_struct_0(X22),X24,X25),u1_pre_topc(X22))))|~v2_pre_topc(X22)|~l1_pre_topc(X22))))&(((m1_subset_1(esk5_1(X22),k1_zfmisc_1(u1_struct_0(X22)))|m1_subset_1(esk4_1(X22),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X22))))|~r2_hidden(u1_struct_0(X22),u1_pre_topc(X22))|v2_pre_topc(X22)|~l1_pre_topc(X22))&((m1_subset_1(esk6_1(X22),k1_zfmisc_1(u1_struct_0(X22)))|m1_subset_1(esk4_1(X22),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X22))))|~r2_hidden(u1_struct_0(X22),u1_pre_topc(X22))|v2_pre_topc(X22)|~l1_pre_topc(X22))&(((r2_hidden(esk5_1(X22),u1_pre_topc(X22))|m1_subset_1(esk4_1(X22),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X22))))|~r2_hidden(u1_struct_0(X22),u1_pre_topc(X22))|v2_pre_topc(X22)|~l1_pre_topc(X22))&(r2_hidden(esk6_1(X22),u1_pre_topc(X22))|m1_subset_1(esk4_1(X22),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X22))))|~r2_hidden(u1_struct_0(X22),u1_pre_topc(X22))|v2_pre_topc(X22)|~l1_pre_topc(X22)))&(~r2_hidden(k9_subset_1(u1_struct_0(X22),esk5_1(X22),esk6_1(X22)),u1_pre_topc(X22))|m1_subset_1(esk4_1(X22),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X22))))|~r2_hidden(u1_struct_0(X22),u1_pre_topc(X22))|v2_pre_topc(X22)|~l1_pre_topc(X22)))))&(((m1_subset_1(esk5_1(X22),k1_zfmisc_1(u1_struct_0(X22)))|r1_tarski(esk4_1(X22),u1_pre_topc(X22))|~r2_hidden(u1_struct_0(X22),u1_pre_topc(X22))|v2_pre_topc(X22)|~l1_pre_topc(X22))&((m1_subset_1(esk6_1(X22),k1_zfmisc_1(u1_struct_0(X22)))|r1_tarski(esk4_1(X22),u1_pre_topc(X22))|~r2_hidden(u1_struct_0(X22),u1_pre_topc(X22))|v2_pre_topc(X22)|~l1_pre_topc(X22))&(((r2_hidden(esk5_1(X22),u1_pre_topc(X22))|r1_tarski(esk4_1(X22),u1_pre_topc(X22))|~r2_hidden(u1_struct_0(X22),u1_pre_topc(X22))|v2_pre_topc(X22)|~l1_pre_topc(X22))&(r2_hidden(esk6_1(X22),u1_pre_topc(X22))|r1_tarski(esk4_1(X22),u1_pre_topc(X22))|~r2_hidden(u1_struct_0(X22),u1_pre_topc(X22))|v2_pre_topc(X22)|~l1_pre_topc(X22)))&(~r2_hidden(k9_subset_1(u1_struct_0(X22),esk5_1(X22),esk6_1(X22)),u1_pre_topc(X22))|r1_tarski(esk4_1(X22),u1_pre_topc(X22))|~r2_hidden(u1_struct_0(X22),u1_pre_topc(X22))|v2_pre_topc(X22)|~l1_pre_topc(X22)))))&((m1_subset_1(esk5_1(X22),k1_zfmisc_1(u1_struct_0(X22)))|~r2_hidden(k5_setfam_1(u1_struct_0(X22),esk4_1(X22)),u1_pre_topc(X22))|~r2_hidden(u1_struct_0(X22),u1_pre_topc(X22))|v2_pre_topc(X22)|~l1_pre_topc(X22))&((m1_subset_1(esk6_1(X22),k1_zfmisc_1(u1_struct_0(X22)))|~r2_hidden(k5_setfam_1(u1_struct_0(X22),esk4_1(X22)),u1_pre_topc(X22))|~r2_hidden(u1_struct_0(X22),u1_pre_topc(X22))|v2_pre_topc(X22)|~l1_pre_topc(X22))&(((r2_hidden(esk5_1(X22),u1_pre_topc(X22))|~r2_hidden(k5_setfam_1(u1_struct_0(X22),esk4_1(X22)),u1_pre_topc(X22))|~r2_hidden(u1_struct_0(X22),u1_pre_topc(X22))|v2_pre_topc(X22)|~l1_pre_topc(X22))&(r2_hidden(esk6_1(X22),u1_pre_topc(X22))|~r2_hidden(k5_setfam_1(u1_struct_0(X22),esk4_1(X22)),u1_pre_topc(X22))|~r2_hidden(u1_struct_0(X22),u1_pre_topc(X22))|v2_pre_topc(X22)|~l1_pre_topc(X22)))&(~r2_hidden(k9_subset_1(u1_struct_0(X22),esk5_1(X22),esk6_1(X22)),u1_pre_topc(X22))|~r2_hidden(k5_setfam_1(u1_struct_0(X22),esk4_1(X22)),u1_pre_topc(X22))|~r2_hidden(u1_struct_0(X22),u1_pre_topc(X22))|v2_pre_topc(X22)|~l1_pre_topc(X22)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])).
cnf(c_0_50, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_51, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_52, negated_conjecture, (~r2_hidden(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))|~r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_42]), c_0_33])])).
cnf(c_0_53, plain, (r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_54, plain, ![X29]:((m1_subset_1(esk7_1(X29),k1_zfmisc_1(u1_struct_0(X29)))|~l1_pre_topc(X29))&(v2_tops_1(esk7_1(X29),X29)|~l1_pre_topc(X29))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[rc5_tops_1])])])])).
cnf(c_0_55, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_56, negated_conjecture, (~r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_33]), c_0_34])])).
fof(c_0_57, plain, ![X31, X32, X33]:(~r2_hidden(X31,X32)|(~m1_subset_1(X32,k1_zfmisc_1(X33))|~v1_xboole_0(X33))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
cnf(c_0_58, plain, (m1_subset_1(esk7_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_59, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0), inference(sr,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_60, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[c_0_51, c_0_56])).
cnf(c_0_61, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_62, negated_conjecture, (m1_subset_1(esk7_1(esk1_0),k1_zfmisc_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_33])])).
cnf(c_0_63, negated_conjecture, (v1_xboole_0(k1_xboole_0)), inference(rw,[status(thm)],[c_0_60, c_0_59])).
fof(c_0_64, plain, ![X34, X36, X37]:((r2_hidden(esk8_1(X34),X34)|X34=k1_xboole_0)&((~r2_hidden(X36,X34)|esk8_1(X34)!=k4_tarski(X36,X37)|X34=k1_xboole_0)&(~r2_hidden(X37,X34)|esk8_1(X34)!=k4_tarski(X36,X37)|X34=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_mcart_1])])])])])])).
fof(c_0_65, plain, ![X38]:(v2_struct_0(X38)|(~v2_pre_topc(X38)|~l1_pre_topc(X38))|~v2_tops_1(k2_struct_0(X38),X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc14_tops_1])])])).
cnf(c_0_66, negated_conjecture, (~r2_hidden(X1,esk7_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63])])).
cnf(c_0_67, plain, (r2_hidden(esk8_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_68, plain, (v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_tops_1(k2_struct_0(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_69, plain, (v2_tops_1(esk7_1(X1),X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_70, negated_conjecture, (esk7_1(esk1_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_71, plain, (v2_struct_0(X1)|~v2_tops_1(u1_struct_0(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_68, c_0_41])).
cnf(c_0_72, negated_conjecture, (v2_tops_1(k1_xboole_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_33])])).
cnf(c_0_73, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_74, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_59]), c_0_72]), c_0_33]), c_0_34])]), c_0_73]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 75
# Proof object clause steps            : 42
# Proof object formula steps           : 33
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 29
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 46
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 45
# Processed clauses                    : 147
# ...of these trivial                  : 2
# ...subsumed                          : 25
# ...remaining for further processing  : 120
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 13
# Generated clauses                    : 82
# ...of the previous two non-trivial   : 85
# Contextual simplify-reflections      : 0
# Paramodulations                      : 78
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 58
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 42
# Current number of unprocessed clauses: 20
# ...number of literals in the above   : 89
# Current number of archived formulas  : 0
# Current number of archived clauses   : 63
# Clause-clause subsumption calls (NU) : 1547
# Rec. Clause-clause subsumption calls : 396
# Non-unit clause-clause subsumptions  : 8
# Unit Clause-clause subsumption calls : 35
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5012

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.000 s
# Total time               : 0.013 s
# Maximum resident set size: 28612 pages
