# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/6dad296aff1e4fcf9ba5a7dde0594a3f.E_input', d4_xboole_0)).
fof(d6_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:k2_wellord1(X1,X2)=k3_xboole_0(X1,k2_zfmisc_1(X2,X2))), file('data/MPTP/proofs/6dad296aff1e4fcf9ba5a7dde0594a3f.E_input', d6_wellord1)).
fof(l3_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v4_relat_2(X1)<=>![X2, X3]:((r2_hidden(k4_tarski(X2,X3),X1)&r2_hidden(k4_tarski(X3,X2),X1))=>X2=X3))), file('data/MPTP/proofs/6dad296aff1e4fcf9ba5a7dde0594a3f.E_input', l3_wellord1)).
fof(dt_k2_wellord1, axiom, ![X1, X2]:(v1_relat_1(X1)=>v1_relat_1(k2_wellord1(X1,X2))), file('data/MPTP/proofs/6dad296aff1e4fcf9ba5a7dde0594a3f.E_input', dt_k2_wellord1)).
fof(t25_wellord1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>(v4_relat_2(X2)=>v4_relat_2(k2_wellord1(X2,X1)))), file('data/MPTP/proofs/6dad296aff1e4fcf9ba5a7dde0594a3f.E_input', t25_wellord1)).
fof(c_0_5, plain, ![X5, X6, X7, X8, X9, X10, X11, X12]:((((r2_hidden(X8,X5)|~r2_hidden(X8,X7)|X7!=k3_xboole_0(X5,X6))&(r2_hidden(X8,X6)|~r2_hidden(X8,X7)|X7!=k3_xboole_0(X5,X6)))&(~r2_hidden(X9,X5)|~r2_hidden(X9,X6)|r2_hidden(X9,X7)|X7!=k3_xboole_0(X5,X6)))&((~r2_hidden(esk1_3(X10,X11,X12),X12)|(~r2_hidden(esk1_3(X10,X11,X12),X10)|~r2_hidden(esk1_3(X10,X11,X12),X11))|X12=k3_xboole_0(X10,X11))&((r2_hidden(esk1_3(X10,X11,X12),X10)|r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k3_xboole_0(X10,X11))&(r2_hidden(esk1_3(X10,X11,X12),X11)|r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k3_xboole_0(X10,X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_6, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_7, plain, ![X14, X15]:(~v1_relat_1(X14)|k2_wellord1(X14,X15)=k3_xboole_0(X14,k2_zfmisc_1(X15,X15))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_wellord1])])])).
cnf(c_0_8, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_6])).
cnf(c_0_9, plain, (k2_wellord1(X1,X2)=k3_xboole_0(X1,k2_zfmisc_1(X2,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_10, plain, ![X16, X17, X18]:((~v4_relat_2(X16)|(~r2_hidden(k4_tarski(X17,X18),X16)|~r2_hidden(k4_tarski(X18,X17),X16)|X17=X18)|~v1_relat_1(X16))&(((r2_hidden(k4_tarski(esk2_1(X16),esk3_1(X16)),X16)|v4_relat_2(X16)|~v1_relat_1(X16))&(r2_hidden(k4_tarski(esk3_1(X16),esk2_1(X16)),X16)|v4_relat_2(X16)|~v1_relat_1(X16)))&(esk2_1(X16)!=esk3_1(X16)|v4_relat_2(X16)|~v1_relat_1(X16)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_wellord1])])])])])).
fof(c_0_11, plain, ![X21, X22]:(~v1_relat_1(X21)|v1_relat_1(k2_wellord1(X21,X22))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_wellord1])])).
cnf(c_0_12, plain, (r2_hidden(X1,X2)|~v1_relat_1(X2)|~r2_hidden(X1,k2_wellord1(X2,X3))), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_13, plain, (r2_hidden(k4_tarski(esk2_1(X1),esk3_1(X1)),X1)|v4_relat_2(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (v1_relat_1(k2_wellord1(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (X2=X3|~v4_relat_2(X1)|~r2_hidden(k4_tarski(X2,X3),X1)|~r2_hidden(k4_tarski(X3,X2),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (v4_relat_2(k2_wellord1(X1,X2))|r2_hidden(k4_tarski(esk2_1(k2_wellord1(X1,X2)),esk3_1(k2_wellord1(X1,X2))),X1)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14])).
cnf(c_0_17, plain, (r2_hidden(k4_tarski(esk3_1(X1),esk2_1(X1)),X1)|v4_relat_2(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_18, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>(v4_relat_2(X2)=>v4_relat_2(k2_wellord1(X2,X1))))), inference(assume_negation,[status(cth)],[t25_wellord1])).
cnf(c_0_19, plain, (esk3_1(k2_wellord1(X1,X2))=esk2_1(k2_wellord1(X1,X2))|v4_relat_2(k2_wellord1(X1,X2))|~v4_relat_2(X1)|~v1_relat_1(X1)|~r2_hidden(k4_tarski(esk3_1(k2_wellord1(X1,X2)),esk2_1(k2_wellord1(X1,X2))),X1)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_20, plain, (v4_relat_2(k2_wellord1(X1,X2))|r2_hidden(k4_tarski(esk3_1(k2_wellord1(X1,X2)),esk2_1(k2_wellord1(X1,X2))),X1)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_17]), c_0_14])).
fof(c_0_21, negated_conjecture, (v1_relat_1(esk5_0)&(v4_relat_2(esk5_0)&~v4_relat_2(k2_wellord1(esk5_0,esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).
cnf(c_0_22, plain, (v4_relat_2(X1)|esk2_1(X1)!=esk3_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_23, plain, (esk3_1(k2_wellord1(X1,X2))=esk2_1(k2_wellord1(X1,X2))|v4_relat_2(k2_wellord1(X1,X2))|~v4_relat_2(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_24, negated_conjecture, (~v4_relat_2(k2_wellord1(esk5_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_25, plain, (v4_relat_2(k2_wellord1(X1,X2))|~v4_relat_2(X1)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_14])).
cnf(c_0_26, negated_conjecture, (v4_relat_2(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, negated_conjecture, (v1_relat_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26]), c_0_27])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 29
# Proof object clause steps            : 18
# Proof object formula steps           : 11
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 7
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 15
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 15
# Processed clauses                    : 52
# ...of these trivial                  : 0
# ...subsumed                          : 4
# ...remaining for further processing  : 48
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 0
# Generated clauses                    : 110
# ...of the previous two non-trivial   : 88
# Contextual simplify-reflections      : 5
# Paramodulations                      : 103
# Factorizations                       : 4
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 28
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 24
# Current number of unprocessed clauses: 66
# ...number of literals in the above   : 255
# Current number of archived formulas  : 0
# Current number of archived clauses   : 17
# Clause-clause subsumption calls (NU) : 149
# Rec. Clause-clause subsumption calls : 87
# Non-unit clause-clause subsumptions  : 11
# Unit Clause-clause subsumption calls : 7
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2827

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28712 pages
