# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t32_ordinal1, conjecture, ![X1, X2]:(v3_ordinal1(X2)=>~((r1_tarski(X1,X2)&(~(X1=k1_xboole_0)&![X3]:(v3_ordinal1(X3)=>~((r2_hidden(X3,X1)&![X4]:(v3_ordinal1(X4)=>(r2_hidden(X4,X1)=>r1_ordinal1(X3,X4)))))))))), file('data/MPTP/proofs/6f173df7e74e4a94a37f9bfd682bb738.E_input', t32_ordinal1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/6f173df7e74e4a94a37f9bfd682bb738.E_input', d3_tarski)).
fof(t26_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r1_ordinal1(X1,X2)|r2_hidden(X2,X1)))), file('data/MPTP/proofs/6f173df7e74e4a94a37f9bfd682bb738.E_input', t26_ordinal1)).
fof(t23_ordinal1, axiom, ![X1, X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)=>v3_ordinal1(X1))), file('data/MPTP/proofs/6f173df7e74e4a94a37f9bfd682bb738.E_input', t23_ordinal1)).
fof(t3_ordinal1, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(r2_hidden(X2,X3)&r2_hidden(X3,X1)))), file('data/MPTP/proofs/6f173df7e74e4a94a37f9bfd682bb738.E_input', t3_ordinal1)).
fof(t136_zfmisc_1, axiom, ![X1]:?[X2]:(r2_hidden(X1,X2)&(![X3, X4]:((r2_hidden(X3,X2)&r1_tarski(X4,X3))=>r2_hidden(X4,X2))&(![X3]:(r2_hidden(X3,X2)=>r2_hidden(k1_zfmisc_1(X3),X2))&![X3]:~((r1_tarski(X3,X2)&(~(r2_tarski(X3,X2))&~(r2_hidden(X3,X2)))))))), file('data/MPTP/proofs/6f173df7e74e4a94a37f9bfd682bb738.E_input', t136_zfmisc_1)).
fof(t7_tarski, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&r2_hidden(X4,X3))))))), file('data/MPTP/proofs/6f173df7e74e4a94a37f9bfd682bb738.E_input', t7_tarski)).
fof(t6_setfam_1, axiom, ![X1, X2]:(![X3]:(r2_hidden(X3,X1)=>r1_tarski(X2,X3))=>(X1=k1_xboole_0|r1_tarski(X2,k1_setfam_1(X1)))), file('data/MPTP/proofs/6f173df7e74e4a94a37f9bfd682bb738.E_input', t6_setfam_1)).
fof(c_0_8, negated_conjecture, ~(![X1, X2]:(v3_ordinal1(X2)=>~((r1_tarski(X1,X2)&(~(X1=k1_xboole_0)&![X3]:(v3_ordinal1(X3)=>~((r2_hidden(X3,X1)&![X4]:(v3_ordinal1(X4)=>(r2_hidden(X4,X1)=>r1_ordinal1(X3,X4))))))))))), inference(assume_negation,[status(cth)],[t32_ordinal1])).
fof(c_0_9, plain, ![X9, X10, X11, X12, X13]:((~r1_tarski(X9,X10)|(~r2_hidden(X11,X9)|r2_hidden(X11,X10)))&((r2_hidden(esk4_2(X12,X13),X12)|r1_tarski(X12,X13))&(~r2_hidden(esk4_2(X12,X13),X13)|r1_tarski(X12,X13)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_10, negated_conjecture, ![X7]:(v3_ordinal1(esk2_0)&(r1_tarski(esk1_0,esk2_0)&(esk1_0!=k1_xboole_0&((v3_ordinal1(esk3_1(X7))|~r2_hidden(X7,esk1_0)|~v3_ordinal1(X7))&((r2_hidden(esk3_1(X7),esk1_0)|~r2_hidden(X7,esk1_0)|~v3_ordinal1(X7))&(~r1_ordinal1(X7,esk3_1(X7))|~r2_hidden(X7,esk1_0)|~v3_ordinal1(X7))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])])).
fof(c_0_11, plain, ![X15, X16]:(~v3_ordinal1(X15)|(~v3_ordinal1(X16)|(r1_ordinal1(X15,X16)|r2_hidden(X16,X15)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t26_ordinal1])])])).
fof(c_0_12, plain, ![X17, X18]:(~v3_ordinal1(X18)|(~r2_hidden(X17,X18)|v3_ordinal1(X17))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_ordinal1])])).
cnf(c_0_13, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, plain, ![X26, X27, X28]:(~r2_hidden(X26,X27)|(~r2_hidden(X27,X28)|~r2_hidden(X28,X26))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_ordinal1])])).
fof(c_0_16, plain, ![X29, X31, X32, X33, X34]:(r2_hidden(X29,esk7_1(X29))&((~r2_hidden(X31,esk7_1(X29))|~r1_tarski(X32,X31)|r2_hidden(X32,esk7_1(X29)))&((~r2_hidden(X33,esk7_1(X29))|r2_hidden(k1_zfmisc_1(X33),esk7_1(X29)))&(~r1_tarski(X34,esk7_1(X29))|(r2_tarski(X34,esk7_1(X29))|r2_hidden(X34,esk7_1(X29))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t136_zfmisc_1])])])])])).
fof(c_0_17, plain, ![X19, X20, X22]:((r2_hidden(esk5_2(X19,X20),X20)|~r2_hidden(X19,X20))&(~r2_hidden(X22,X20)|~r2_hidden(X22,esk5_2(X19,X20))|~r2_hidden(X19,X20))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tarski])])])])])).
cnf(c_0_18, negated_conjecture, (~r1_ordinal1(X1,esk3_1(X1))|~r2_hidden(X1,esk1_0)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19, plain, (r1_ordinal1(X1,X2)|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20, negated_conjecture, (v3_ordinal1(esk3_1(X1))|~r2_hidden(X1,esk1_0)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_21, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_23, negated_conjecture, (v3_ordinal1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_24, plain, (~r2_hidden(X1,X2)|~r2_hidden(X2,X3)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25, plain, (r2_hidden(X1,esk7_1(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26, plain, (r2_hidden(X3,esk7_1(X2))|~r2_hidden(X1,esk7_1(X2))|~r1_tarski(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_27, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,esk5_2(X3,X2))|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28, negated_conjecture, (r2_hidden(esk3_1(X1),X1)|~r2_hidden(X1,esk1_0)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])).
cnf(c_0_29, negated_conjecture, (v3_ordinal1(X1)|~r2_hidden(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])])).
cnf(c_0_30, plain, (~r2_hidden(esk7_1(X1),X2)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31, plain, (r2_hidden(X1,esk7_1(X2))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_26, c_0_25])).
fof(c_0_32, plain, ![X23, X24]:((r2_hidden(esk6_2(X23,X24),X23)|(X23=k1_xboole_0|r1_tarski(X24,k1_setfam_1(X23))))&(~r1_tarski(X24,esk6_2(X23,X24))|(X23=k1_xboole_0|r1_tarski(X24,k1_setfam_1(X23))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_setfam_1])])])])).
cnf(c_0_33, negated_conjecture, (~r2_hidden(esk3_1(esk5_2(X1,X2)),X2)|~r2_hidden(esk5_2(X1,X2),esk1_0)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])).
cnf(c_0_34, negated_conjecture, (r2_hidden(esk3_1(X1),esk1_0)|~r2_hidden(X1,esk1_0)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_35, plain, (r2_hidden(esk5_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_36, plain, (~r2_hidden(esk7_1(X1),X2)|~r1_tarski(esk7_1(X2),X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_37, plain, (r2_hidden(esk6_2(X1,X2),X1)|X1=k1_xboole_0|r1_tarski(X2,k1_setfam_1(X1))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, negated_conjecture, (~r2_hidden(X1,esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_29]), c_0_35])).
cnf(c_0_39, plain, (X1=k1_xboole_0|r2_hidden(esk6_2(X1,esk7_1(X2)),X1)|~r2_hidden(esk7_1(k1_setfam_1(X1)),X2)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_40, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_41, negated_conjecture, (~r2_hidden(esk7_1(k1_setfam_1(esk1_0)),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])).
cnf(c_0_42, negated_conjecture, ($false), inference(spm,[status(thm)],[c_0_41, c_0_25]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 43
# Proof object clause steps            : 26
# Proof object formula steps           : 17
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 7
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 20
# Processed clauses                    : 254
# ...of these trivial                  : 0
# ...subsumed                          : 72
# ...remaining for further processing  : 182
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 17
# Backward-rewritten                   : 2
# Generated clauses                    : 351
# ...of the previous two non-trivial   : 344
# Contextual simplify-reflections      : 11
# Paramodulations                      : 351
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 143
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 130
# Current number of unprocessed clauses: 117
# ...number of literals in the above   : 481
# Current number of archived formulas  : 0
# Current number of archived clauses   : 39
# Clause-clause subsumption calls (NU) : 1799
# Rec. Clause-clause subsumption calls : 1139
# Non-unit clause-clause subsumptions  : 54
# Unit Clause-clause subsumption calls : 82
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7427

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.006 s
# Total time               : 0.015 s
# Maximum resident set size: 28668 pages
