# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t39_orders_2, conjecture, ![X1]:((v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(?[X4]:((v6_orders_2(X4,X1)&m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1))))&(r2_hidden(X2,X4)&r2_hidden(X3,X4)))<=>(r2_orders_2(X1,X2,X3)<=>~(r1_orders_2(X1,X3,X2))))))), file('data/MPTP/proofs/6fc830227c3b4465a5e04480ef8bfd35.E_input', t39_orders_2)).
fof(t38_orders_2, axiom, ![X1]:((v3_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(~((?[X4]:((v6_orders_2(X4,X1)&m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1))))&(r2_hidden(X2,X4)&r2_hidden(X3,X4)))&(~(r1_orders_2(X1,X2,X3))&~(r1_orders_2(X1,X3,X2)))))&~(((r1_orders_2(X1,X2,X3)|r1_orders_2(X1,X3,X2))&![X4]:((v6_orders_2(X4,X1)&m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1))))=>~((r2_hidden(X2,X4)&r2_hidden(X3,X4)))))))))), file('data/MPTP/proofs/6fc830227c3b4465a5e04480ef8bfd35.E_input', t38_orders_2)).
fof(d10_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_orders_2(X1,X2,X3)<=>(r1_orders_2(X1,X2,X3)&~(X2=X3)))))), file('data/MPTP/proofs/6fc830227c3b4465a5e04480ef8bfd35.E_input', d10_orders_2)).
fof(t30_orders_2, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~((r1_orders_2(X1,X2,X3)&r2_orders_2(X1,X3,X2)))))), file('data/MPTP/proofs/6fc830227c3b4465a5e04480ef8bfd35.E_input', t30_orders_2)).
fof(c_0_4, negated_conjecture, ~(![X1]:((v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(?[X4]:((v6_orders_2(X4,X1)&m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1))))&(r2_hidden(X2,X4)&r2_hidden(X3,X4)))<=>(r2_orders_2(X1,X2,X3)<=>~(r1_orders_2(X1,X3,X2)))))))), inference(assume_negation,[status(cth)],[t39_orders_2])).
fof(c_0_5, plain, ![X10, X11, X12, X13]:((~v6_orders_2(X13,X10)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X10)))|(~r2_hidden(X11,X13)|~r2_hidden(X12,X13))|(r1_orders_2(X10,X11,X12)|r1_orders_2(X10,X12,X11))|~m1_subset_1(X12,u1_struct_0(X10))|~m1_subset_1(X11,u1_struct_0(X10))|(~v3_orders_2(X10)|~l1_orders_2(X10)))&((((v6_orders_2(esk5_3(X10,X11,X12),X10)|~r1_orders_2(X10,X11,X12)|~m1_subset_1(X12,u1_struct_0(X10))|~m1_subset_1(X11,u1_struct_0(X10))|(~v3_orders_2(X10)|~l1_orders_2(X10)))&(m1_subset_1(esk5_3(X10,X11,X12),k1_zfmisc_1(u1_struct_0(X10)))|~r1_orders_2(X10,X11,X12)|~m1_subset_1(X12,u1_struct_0(X10))|~m1_subset_1(X11,u1_struct_0(X10))|(~v3_orders_2(X10)|~l1_orders_2(X10))))&((r2_hidden(X11,esk5_3(X10,X11,X12))|~r1_orders_2(X10,X11,X12)|~m1_subset_1(X12,u1_struct_0(X10))|~m1_subset_1(X11,u1_struct_0(X10))|(~v3_orders_2(X10)|~l1_orders_2(X10)))&(r2_hidden(X12,esk5_3(X10,X11,X12))|~r1_orders_2(X10,X11,X12)|~m1_subset_1(X12,u1_struct_0(X10))|~m1_subset_1(X11,u1_struct_0(X10))|(~v3_orders_2(X10)|~l1_orders_2(X10)))))&(((v6_orders_2(esk5_3(X10,X11,X12),X10)|~r1_orders_2(X10,X12,X11)|~m1_subset_1(X12,u1_struct_0(X10))|~m1_subset_1(X11,u1_struct_0(X10))|(~v3_orders_2(X10)|~l1_orders_2(X10)))&(m1_subset_1(esk5_3(X10,X11,X12),k1_zfmisc_1(u1_struct_0(X10)))|~r1_orders_2(X10,X12,X11)|~m1_subset_1(X12,u1_struct_0(X10))|~m1_subset_1(X11,u1_struct_0(X10))|(~v3_orders_2(X10)|~l1_orders_2(X10))))&((r2_hidden(X11,esk5_3(X10,X11,X12))|~r1_orders_2(X10,X12,X11)|~m1_subset_1(X12,u1_struct_0(X10))|~m1_subset_1(X11,u1_struct_0(X10))|(~v3_orders_2(X10)|~l1_orders_2(X10)))&(r2_hidden(X12,esk5_3(X10,X11,X12))|~r1_orders_2(X10,X12,X11)|~m1_subset_1(X12,u1_struct_0(X10))|~m1_subset_1(X11,u1_struct_0(X10))|(~v3_orders_2(X10)|~l1_orders_2(X10))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t38_orders_2])])])])])])).
fof(c_0_6, negated_conjecture, ![X8]:((v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(((~r2_orders_2(esk1_0,esk2_0,esk3_0)|r1_orders_2(esk1_0,esk3_0,esk2_0)|(~v6_orders_2(X8,esk1_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(esk1_0)))|(~r2_hidden(esk2_0,X8)|~r2_hidden(esk3_0,X8))))&(r2_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk3_0,esk2_0)|(~v6_orders_2(X8,esk1_0)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(esk1_0)))|(~r2_hidden(esk2_0,X8)|~r2_hidden(esk3_0,X8)))))&((((~r2_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk3_0,esk2_0)|v6_orders_2(esk4_0,esk1_0))&(r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_orders_2(esk1_0,esk2_0,esk3_0)|v6_orders_2(esk4_0,esk1_0)))&((~r2_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk3_0,esk2_0)|m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0))))&(r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_orders_2(esk1_0,esk2_0,esk3_0)|m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0))))))&(((~r2_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_hidden(esk2_0,esk4_0))&(r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_orders_2(esk1_0,esk2_0,esk3_0)|r2_hidden(esk2_0,esk4_0)))&((~r2_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_hidden(esk3_0,esk4_0))&(r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_orders_2(esk1_0,esk2_0,esk3_0)|r2_hidden(esk3_0,esk4_0))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_4])])])])])])).
cnf(c_0_7, plain, (r1_orders_2(X2,X3,X4)|r1_orders_2(X2,X4,X3)|~v6_orders_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X3,X1)|~r2_hidden(X4,X1)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~v3_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_orders_2(esk1_0,esk2_0,esk3_0)|m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_orders_2(esk1_0,esk2_0,esk3_0)|v6_orders_2(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_12, plain, ![X15, X16, X17]:(((r1_orders_2(X15,X16,X17)|~r2_orders_2(X15,X16,X17)|~m1_subset_1(X17,u1_struct_0(X15))|~m1_subset_1(X16,u1_struct_0(X15))|~l1_orders_2(X15))&(X16!=X17|~r2_orders_2(X15,X16,X17)|~m1_subset_1(X17,u1_struct_0(X15))|~m1_subset_1(X16,u1_struct_0(X15))|~l1_orders_2(X15)))&(~r1_orders_2(X15,X16,X17)|X16=X17|r2_orders_2(X15,X16,X17)|~m1_subset_1(X17,u1_struct_0(X15))|~m1_subset_1(X16,u1_struct_0(X15))|~l1_orders_2(X15))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_orders_2])])])])])).
cnf(c_0_13, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk2_0)|r1_orders_2(esk1_0,X1,X2)|r1_orders_2(esk1_0,X2,X1)|r2_orders_2(esk1_0,esk2_0,esk3_0)|~r2_hidden(X2,esk4_0)|~r2_hidden(X1,esk4_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_9]), c_0_10])]), c_0_11])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_orders_2(esk1_0,esk2_0,esk3_0)|r2_hidden(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_16, plain, ![X18, X19, X20]:(~v5_orders_2(X18)|~l1_orders_2(X18)|(~m1_subset_1(X19,u1_struct_0(X18))|(~m1_subset_1(X20,u1_struct_0(X18))|(~r1_orders_2(X18,X19,X20)|~r2_orders_2(X18,X20,X19))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_orders_2])])])).
cnf(c_0_17, plain, (r1_orders_2(X1,X2,X3)|~r2_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk2_0)|r1_orders_2(esk1_0,esk3_0,X1)|r1_orders_2(esk1_0,X1,esk3_0)|r2_orders_2(esk1_0,esk2_0,esk3_0)|~r2_hidden(X1,esk4_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])).
cnf(c_0_19, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_20, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_orders_2(esk1_0,esk2_0,esk3_0)|r2_hidden(esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_21, plain, (~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)|~r2_orders_2(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_23, negated_conjecture, (r1_orders_2(esk1_0,X1,esk3_0)|~r2_orders_2(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_14]), c_0_9])])).
cnf(c_0_24, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)|r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_orders_2(esk1_0,esk2_0,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])).
cnf(c_0_25, negated_conjecture, (~r1_orders_2(esk1_0,X1,esk2_0)|~r2_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_19]), c_0_9]), c_0_22])])).
cnf(c_0_26, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk2_0)|r1_orders_2(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_19])])).
cnf(c_0_27, plain, (m1_subset_1(esk5_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_28, plain, (v6_orders_2(esk5_3(X1,X2,X3),X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_29, plain, (r2_hidden(X1,esk5_3(X2,X1,X3))|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~v3_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_30, plain, (r2_hidden(X1,esk5_3(X2,X3,X1))|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~v3_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_31, negated_conjecture, (r2_orders_2(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk3_0,esk2_0)|~v6_orders_2(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(esk2_0,X1)|~r2_hidden(esk3_0,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_32, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)|~r2_orders_2(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_14])])).
cnf(c_0_33, negated_conjecture, (m1_subset_1(esk5_3(esk1_0,X1,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_19]), c_0_9]), c_0_10])])).
cnf(c_0_34, negated_conjecture, (v6_orders_2(esk5_3(esk1_0,X1,esk2_0),esk1_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_19]), c_0_9]), c_0_10])])).
cnf(c_0_35, negated_conjecture, (r2_hidden(X1,esk5_3(esk1_0,X1,esk2_0))|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_19]), c_0_9]), c_0_10])])).
cnf(c_0_36, negated_conjecture, (r2_hidden(X1,esk5_3(esk1_0,esk3_0,X1))|~r1_orders_2(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_14]), c_0_9]), c_0_10])])).
cnf(c_0_37, plain, (X2=X3|r2_orders_2(X1,X2,X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_38, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)|~r2_hidden(esk2_0,X1)|~r2_hidden(esk3_0,X1)|~v6_orders_2(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_26]), c_0_32])).
cnf(c_0_39, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)|m1_subset_1(esk5_3(esk1_0,esk3_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_26]), c_0_14])])).
cnf(c_0_40, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)|v6_orders_2(esk5_3(esk1_0,esk3_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_26]), c_0_14])])).
cnf(c_0_41, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)|r2_hidden(esk3_0,esk5_3(esk1_0,esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_26]), c_0_14])])).
cnf(c_0_42, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)|r2_hidden(esk2_0,esk5_3(esk1_0,esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_26]), c_0_19])])).
cnf(c_0_43, negated_conjecture, (X1=esk3_0|r2_orders_2(esk1_0,X1,esk3_0)|~r1_orders_2(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_14]), c_0_9])])).
cnf(c_0_44, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40]), c_0_41]), c_0_42])).
cnf(c_0_45, plain, (m1_subset_1(esk5_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_46, plain, (r2_hidden(X1,esk5_3(X2,X3,X1))|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~v3_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_47, plain, (r2_hidden(X1,esk5_3(X2,X1,X3))|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~v3_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_48, plain, (v6_orders_2(esk5_3(X1,X2,X3),X1)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_49, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk2_0)|~r2_orders_2(esk1_0,esk2_0,esk3_0)|~v6_orders_2(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(esk2_0,X1)|~r2_hidden(esk3_0,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_50, negated_conjecture, (esk3_0=esk2_0|r2_orders_2(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_19])])).
cnf(c_0_51, negated_conjecture, (m1_subset_1(esk5_3(esk1_0,X1,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_19]), c_0_9]), c_0_10])])).
cnf(c_0_52, negated_conjecture, (r2_hidden(X1,esk5_3(esk1_0,esk3_0,X1))|~r1_orders_2(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_14]), c_0_9]), c_0_10])])).
cnf(c_0_53, negated_conjecture, (r2_hidden(X1,esk5_3(esk1_0,X1,esk2_0))|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_19]), c_0_9]), c_0_10])])).
cnf(c_0_54, negated_conjecture, (v6_orders_2(esk5_3(esk1_0,X1,esk2_0),esk1_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_19]), c_0_9]), c_0_10])])).
cnf(c_0_55, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk3_0)|r1_orders_2(esk1_0,esk3_0,esk2_0)|r2_orders_2(esk1_0,esk2_0,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_14]), c_0_15])).
cnf(c_0_56, negated_conjecture, (m1_subset_1(esk5_3(esk1_0,X1,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_14]), c_0_9]), c_0_10])])).
cnf(c_0_57, negated_conjecture, (r2_hidden(X1,esk5_3(esk1_0,X1,esk3_0))|~r1_orders_2(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_14]), c_0_9]), c_0_10])])).
cnf(c_0_58, negated_conjecture, (r2_hidden(X1,esk5_3(esk1_0,esk2_0,X1))|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_19]), c_0_9]), c_0_10])])).
cnf(c_0_59, negated_conjecture, (v6_orders_2(esk5_3(esk1_0,X1,esk3_0),esk1_0)|~r1_orders_2(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_14]), c_0_9]), c_0_10])])).
cnf(c_0_60, negated_conjecture, (esk3_0=esk2_0|r1_orders_2(esk1_0,esk3_0,esk2_0)|~r2_hidden(esk2_0,X1)|~r2_hidden(esk3_0,X1)|~v6_orders_2(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_61, negated_conjecture, (m1_subset_1(esk5_3(esk1_0,esk3_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_44]), c_0_14])])).
cnf(c_0_62, negated_conjecture, (r2_hidden(esk2_0,esk5_3(esk1_0,esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_44]), c_0_19])])).
cnf(c_0_63, negated_conjecture, (r2_hidden(esk3_0,esk5_3(esk1_0,esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_44]), c_0_14])])).
cnf(c_0_64, negated_conjecture, (v6_orders_2(esk5_3(esk1_0,esk3_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_44]), c_0_14])])).
cnf(c_0_65, negated_conjecture, (~r1_orders_2(esk1_0,X1,esk3_0)|~r2_orders_2(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_14]), c_0_9]), c_0_22])])).
cnf(c_0_66, plain, (X1!=X2|~r2_orders_2(X3,X1,X2)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_67, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk3_0)|r1_orders_2(esk1_0,esk3_0,esk2_0)|~r2_hidden(esk2_0,X1)|~r2_hidden(esk3_0,X1)|~v6_orders_2(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_49, c_0_55])).
cnf(c_0_68, negated_conjecture, (m1_subset_1(esk5_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_44]), c_0_19])])).
cnf(c_0_69, negated_conjecture, (r2_hidden(esk2_0,esk5_3(esk1_0,esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_44]), c_0_19])])).
cnf(c_0_70, negated_conjecture, (r2_hidden(esk3_0,esk5_3(esk1_0,esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_44]), c_0_14])])).
cnf(c_0_71, negated_conjecture, (v6_orders_2(esk5_3(esk1_0,esk2_0,esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_44]), c_0_19])])).
cnf(c_0_72, negated_conjecture, (X1=esk2_0|r2_orders_2(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_19]), c_0_9])])).
cnf(c_0_73, negated_conjecture, (esk3_0=esk2_0|r1_orders_2(esk1_0,esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62]), c_0_63]), c_0_64])])).
cnf(c_0_74, negated_conjecture, (~r2_orders_2(esk1_0,esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_44]), c_0_19])])).
cnf(c_0_75, plain, (~r2_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_66])).
cnf(c_0_76, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk2_0)|r1_orders_2(esk1_0,esk3_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_69]), c_0_70]), c_0_71])])).
cnf(c_0_77, negated_conjecture, (esk3_0=esk2_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_14])]), c_0_74])).
cnf(c_0_78, negated_conjecture, (~r2_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_19]), c_0_9])])).
cnf(c_0_79, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_77]), c_0_77]), c_0_77])])).
cnf(c_0_80, negated_conjecture, (~r2_hidden(esk2_0,X1)|~v6_orders_2(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_77]), c_0_77]), c_0_77])]), c_0_78]), c_0_79])])).
cnf(c_0_81, negated_conjecture, (m1_subset_1(esk5_3(esk1_0,esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(rw,[status(thm)],[c_0_61, c_0_77])).
cnf(c_0_82, negated_conjecture, (r2_hidden(esk2_0,esk5_3(esk1_0,esk2_0,esk2_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_77]), c_0_77])).
cnf(c_0_83, negated_conjecture, (v6_orders_2(esk5_3(esk1_0,esk2_0,esk2_0),esk1_0)), inference(rw,[status(thm)],[c_0_64, c_0_77])).
cnf(c_0_84, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_82]), c_0_83])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 85
# Proof object clause steps            : 76
# Proof object formula steps           : 9
# Proof object conjectures             : 65
# Proof object clause conjectures      : 62
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 46
# Proof object simplifying inferences  : 123
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 28
# Processed clauses                    : 173
# ...of these trivial                  : 9
# ...subsumed                          : 31
# ...remaining for further processing  : 133
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 13
# Backward-rewritten                   : 56
# Generated clauses                    : 167
# ...of the previous two non-trivial   : 186
# Contextual simplify-reflections      : 10
# Paramodulations                      : 166
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 35
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 25
# Current number of unprocessed clauses: 42
# ...number of literals in the above   : 119
# Current number of archived formulas  : 0
# Current number of archived clauses   : 97
# Clause-clause subsumption calls (NU) : 982
# Rec. Clause-clause subsumption calls : 179
# Non-unit clause-clause subsumptions  : 44
# Unit Clause-clause subsumption calls : 119
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6581

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.000 s
# Total time               : 0.013 s
# Maximum resident set size: 28744 pages
