# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2m
# and selection function SelectCQArNTNpEqFirst.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t26_yellow_1, axiom, ![X1]:((v1_relat_1(X1)&(v4_relat_1(X1,k1_xboole_0)&(v1_funct_1(X1)&(v1_partfun1(X1,k1_xboole_0)&v1_yellow_1(X1)))))=>k5_yellow_1(k1_xboole_0,X1)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))), file('data/MPTP/proofs/70d7834cca0f443ba8d8a78905ddc947.E_input', t26_yellow_1)).
fof(fc20_funcop_1, axiom, ![X1, X2]:(v1_relat_1(k2_funcop_1(X1,X2))&(v4_relat_1(k2_funcop_1(X1,X2),X1)&(v1_funct_1(k2_funcop_1(X1,X2))&v1_partfun1(k2_funcop_1(X1,X2),X1)))), file('data/MPTP/proofs/70d7834cca0f443ba8d8a78905ddc947.E_input', fc20_funcop_1)).
fof(t27_yellow_1, conjecture, ![X1]:(l1_orders_2(X1)=>k6_yellow_1(k1_xboole_0,X1)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))), file('data/MPTP/proofs/70d7834cca0f443ba8d8a78905ddc947.E_input', t27_yellow_1)).
fof(fc10_yellow_1, axiom, ![X1, X2]:(l1_orders_2(X2)=>v1_yellow_1(k2_funcop_1(X1,X2))), file('data/MPTP/proofs/70d7834cca0f443ba8d8a78905ddc947.E_input', fc10_yellow_1)).
fof(d5_yellow_1, axiom, ![X1, X2]:(l1_orders_2(X2)=>k6_yellow_1(X1,X2)=k5_yellow_1(X1,k7_funcop_1(X1,X2))), file('data/MPTP/proofs/70d7834cca0f443ba8d8a78905ddc947.E_input', d5_yellow_1)).
fof(redefinition_k7_funcop_1, axiom, ![X1, X2]:k7_funcop_1(X1,X2)=k2_funcop_1(X1,X2), file('data/MPTP/proofs/70d7834cca0f443ba8d8a78905ddc947.E_input', redefinition_k7_funcop_1)).
fof(c_0_6, plain, ![X12]:(~v1_relat_1(X12)|(~v4_relat_1(X12,k1_xboole_0)|(~v1_funct_1(X12)|(~v1_partfun1(X12,k1_xboole_0)|~v1_yellow_1(X12))))|k5_yellow_1(k1_xboole_0,X12)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t26_yellow_1])])).
fof(c_0_7, plain, ![X4, X5]:(v1_relat_1(k2_funcop_1(X4,X5))&(v4_relat_1(k2_funcop_1(X4,X5),X4)&(v1_funct_1(k2_funcop_1(X4,X5))&v1_partfun1(k2_funcop_1(X4,X5),X4)))), inference(variable_rename,[status(thm)],[fc20_funcop_1])).
fof(c_0_8, negated_conjecture, ~(![X1]:(l1_orders_2(X1)=>k6_yellow_1(k1_xboole_0,X1)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0))))), inference(assume_negation,[status(cth)],[t27_yellow_1])).
cnf(c_0_9, plain, (k5_yellow_1(k1_xboole_0,X1)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))|~v1_relat_1(X1)|~v4_relat_1(X1,k1_xboole_0)|~v1_funct_1(X1)|~v1_partfun1(X1,k1_xboole_0)|~v1_yellow_1(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (v1_partfun1(k2_funcop_1(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (v1_funct_1(k2_funcop_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (v4_relat_1(k2_funcop_1(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, plain, (v1_relat_1(k2_funcop_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_14, plain, ![X8, X9]:(~l1_orders_2(X9)|v1_yellow_1(k2_funcop_1(X8,X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc10_yellow_1])])).
fof(c_0_15, negated_conjecture, (l1_orders_2(esk1_0)&k6_yellow_1(k1_xboole_0,esk1_0)!=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
cnf(c_0_16, plain, (k5_yellow_1(k1_xboole_0,k2_funcop_1(k1_xboole_0,X1))=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))|~v1_yellow_1(k2_funcop_1(k1_xboole_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11]), c_0_12]), c_0_13])])).
cnf(c_0_17, plain, (v1_yellow_1(k2_funcop_1(X2,X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, plain, ![X6, X7]:(~l1_orders_2(X7)|k6_yellow_1(X6,X7)=k5_yellow_1(X6,k7_funcop_1(X6,X7))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_yellow_1])])).
fof(c_0_19, plain, ![X10, X11]:k7_funcop_1(X10,X11)=k2_funcop_1(X10,X11), inference(variable_rename,[status(thm)],[redefinition_k7_funcop_1])).
cnf(c_0_20, negated_conjecture, (k6_yellow_1(k1_xboole_0,esk1_0)!=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (k5_yellow_1(k1_xboole_0,k2_funcop_1(k1_xboole_0,X1))=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_22, plain, (k6_yellow_1(X2,X1)=k5_yellow_1(X2,k7_funcop_1(X2,X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, plain, (k7_funcop_1(X1,X2)=k2_funcop_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, negated_conjecture, (k5_yellow_1(k1_xboole_0,k2_funcop_1(k1_xboole_0,X1))!=k6_yellow_1(k1_xboole_0,esk1_0)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_25, plain, (k6_yellow_1(X2,X1)=k5_yellow_1(X2,k2_funcop_1(X2,X1))|~l1_orders_2(X1)), inference(rw,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_26, negated_conjecture, (k6_yellow_1(k1_xboole_0,X1)!=k6_yellow_1(k1_xboole_0,esk1_0)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_27, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_28, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_26]), c_0_27])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 29
# Proof object clause steps            : 16
# Proof object formula steps           : 13
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 7
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 10
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 9
# Processed clauses                    : 22
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 22
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 8
# ...of the previous two non-trivial   : 7
# Contextual simplify-reflections      : 0
# Paramodulations                      : 7
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 13
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 7
# Current number of unprocessed clauses: 3
# ...number of literals in the above   : 7
# Current number of archived formulas  : 0
# Current number of archived clauses   : 10
# Clause-clause subsumption calls (NU) : 4
# Rec. Clause-clause subsumption calls : 4
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 743

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.000 s
# Total time               : 0.009 s
# Maximum resident set size: 28856 pages
