# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', t12_setfam_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', t70_enumset1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', t100_xboole_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', d5_xboole_0)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', t17_xboole_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', idempotence_k3_xboole_0)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', t3_boole)).
fof(t3_xboole_1, axiom, ![X1]:(r1_tarski(X1,k1_xboole_0)=>X1=k1_xboole_0), file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', t3_xboole_1)).
fof(t178_relat_1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>r1_tarski(k10_relat_1(X3,X1),k10_relat_1(X3,X2)))), file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', t178_relat_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', commutativity_k2_tarski)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', t1_xboole_1)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', t83_xboole_1)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', d4_xboole_0)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', t36_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', d10_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', t28_xboole_1)).
fof(t85_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_xboole_0(X1,k4_xboole_0(X3,X2))), file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', t85_xboole_1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', l51_zfmisc_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', t7_xboole_1)).
fof(t175_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>k10_relat_1(X3,k2_xboole_0(X1,X2))=k2_xboole_0(k10_relat_1(X3,X1),k10_relat_1(X3,X2))), file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', t175_relat_1)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', t39_xboole_1)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/7134f8765e3548618be9a8ab56c762ea.E_input', t1_boole)).
fof(c_0_22, plain, ![X14, X15]:k1_setfam_1(k2_tarski(X14,X15))=k3_xboole_0(X14,X15), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_23, plain, ![X28, X29]:k1_enumset1(X28,X28,X29)=k2_tarski(X28,X29), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_24, plain, ![X25, X26]:k4_xboole_0(X25,X26)=k5_xboole_0(X25,k3_xboole_0(X25,X26)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_25, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_27, plain, ![X36, X37, X38, X39, X40, X41, X42, X43]:((((r2_hidden(X39,X36)|~r2_hidden(X39,X38)|X38!=k4_xboole_0(X36,X37))&(~r2_hidden(X39,X37)|~r2_hidden(X39,X38)|X38!=k4_xboole_0(X36,X37)))&(~r2_hidden(X40,X36)|r2_hidden(X40,X37)|r2_hidden(X40,X38)|X38!=k4_xboole_0(X36,X37)))&((~r2_hidden(esk4_3(X41,X42,X43),X43)|(~r2_hidden(esk4_3(X41,X42,X43),X41)|r2_hidden(esk4_3(X41,X42,X43),X42))|X43=k4_xboole_0(X41,X42))&((r2_hidden(esk4_3(X41,X42,X43),X41)|r2_hidden(esk4_3(X41,X42,X43),X43)|X43=k4_xboole_0(X41,X42))&(~r2_hidden(esk4_3(X41,X42,X43),X42)|r2_hidden(esk4_3(X41,X42,X43),X43)|X43=k4_xboole_0(X41,X42))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_28, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_30, plain, ![X45, X46]:r1_tarski(k3_xboole_0(X45,X46),X45), inference(variable_rename,[status(thm)],[t17_xboole_1])).
cnf(c_0_31, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_32, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[c_0_28, c_0_29])).
fof(c_0_33, plain, ![X60]:k3_xboole_0(X60,X60)=X60, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
fof(c_0_34, plain, ![X27]:k4_xboole_0(X27,k1_xboole_0)=X27, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_35, plain, ![X47]:(~r1_tarski(X47,k1_xboole_0)|X47=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_xboole_1])])).
cnf(c_0_36, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_37, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>r1_tarski(k10_relat_1(X3,X1),k10_relat_1(X3,X2))))), inference(assume_negation,[status(cth)],[t178_relat_1])).
fof(c_0_38, plain, ![X61, X62]:k2_tarski(X61,X62)=k2_tarski(X62,X61), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_39, plain, (X3!=k5_xboole_0(X4,k1_setfam_1(k1_enumset1(X4,X4,X2)))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_40, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_41, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_43, plain, (r1_tarski(k1_setfam_1(k1_enumset1(X1,X1,X2)),X1)), inference(rw,[status(thm)],[c_0_36, c_0_29])).
fof(c_0_44, plain, ![X18, X19, X20]:(~r1_tarski(X18,X19)|~r1_tarski(X19,X20)|r1_tarski(X18,X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_45, negated_conjecture, (v1_relat_1(esk3_0)&(r1_tarski(esk1_0,esk2_0)&~r1_tarski(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_37])])])).
cnf(c_0_46, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_47, plain, ![X21, X22]:((~r1_xboole_0(X21,X22)|k4_xboole_0(X21,X22)=X21)&(k4_xboole_0(X21,X22)!=X21|r1_xboole_0(X21,X22))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
cnf(c_0_48, plain, (~r2_hidden(X1,k5_xboole_0(X2,k1_setfam_1(k1_enumset1(X2,X2,X3))))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_39])).
cnf(c_0_49, plain, (k1_setfam_1(k1_enumset1(X1,X1,X1))=X1), inference(rw,[status(thm)],[c_0_40, c_0_29])).
cnf(c_0_50, plain, (k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,k1_xboole_0)))=X1), inference(rw,[status(thm)],[c_0_41, c_0_32])).
cnf(c_0_51, plain, (k1_setfam_1(k1_enumset1(k1_xboole_0,k1_xboole_0,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_42, c_0_43])).
fof(c_0_52, plain, ![X48, X49, X50, X51, X52, X53, X54, X55]:((((r2_hidden(X51,X48)|~r2_hidden(X51,X50)|X50!=k3_xboole_0(X48,X49))&(r2_hidden(X51,X49)|~r2_hidden(X51,X50)|X50!=k3_xboole_0(X48,X49)))&(~r2_hidden(X52,X48)|~r2_hidden(X52,X49)|r2_hidden(X52,X50)|X50!=k3_xboole_0(X48,X49)))&((~r2_hidden(esk5_3(X53,X54,X55),X55)|(~r2_hidden(esk5_3(X53,X54,X55),X53)|~r2_hidden(esk5_3(X53,X54,X55),X54))|X55=k3_xboole_0(X53,X54))&((r2_hidden(esk5_3(X53,X54,X55),X53)|r2_hidden(esk5_3(X53,X54,X55),X55)|X55=k3_xboole_0(X53,X54))&(r2_hidden(esk5_3(X53,X54,X55),X54)|r2_hidden(esk5_3(X53,X54,X55),X55)|X55=k3_xboole_0(X53,X54))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_53, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_54, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_55, plain, ![X10, X11]:r1_tarski(k4_xboole_0(X10,X11),X10), inference(variable_rename,[status(thm)],[t36_xboole_1])).
cnf(c_0_56, plain, (k1_enumset1(X1,X1,X2)=k1_enumset1(X2,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_26]), c_0_26])).
cnf(c_0_57, plain, (k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_58, plain, (~r2_hidden(X1,k5_xboole_0(X2,X2))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_59, plain, (k5_xboole_0(k1_xboole_0,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_60, plain, (r2_hidden(esk5_3(X1,X2,X3),X1)|r2_hidden(esk5_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_61, negated_conjecture, (r1_tarski(X1,esk2_0)|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_62, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
fof(c_0_63, plain, ![X12, X13]:(((r1_tarski(X12,X13)|X12!=X13)&(r1_tarski(X13,X12)|X12!=X13))&(~r1_tarski(X12,X13)|~r1_tarski(X13,X12)|X12=X13)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_64, plain, ![X16, X17]:(~r1_tarski(X16,X17)|k3_xboole_0(X16,X17)=X16), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_65, plain, (~r2_hidden(X1,k5_xboole_0(X2,k1_setfam_1(k1_enumset1(X3,X3,X2))))|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_48, c_0_56])).
cnf(c_0_66, plain, (k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X2)))=X1|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_57, c_0_32])).
cnf(c_0_67, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_68, plain, (X3=k1_setfam_1(k1_enumset1(X1,X1,X2))|r2_hidden(esk5_3(X1,X2,X3),X3)|r2_hidden(esk5_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_60, c_0_29])).
fof(c_0_69, plain, ![X30, X31, X32]:(~r1_tarski(X30,X31)|r1_xboole_0(X30,k4_xboole_0(X32,X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t85_xboole_1])])).
cnf(c_0_70, negated_conjecture, (r1_tarski(X1,esk2_0)|~r1_tarski(X2,esk1_0)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_53, c_0_61])).
cnf(c_0_71, plain, (r1_tarski(k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X2))),X1)), inference(rw,[status(thm)],[c_0_62, c_0_32])).
cnf(c_0_72, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_73, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_64])).
fof(c_0_74, plain, ![X23, X24]:k3_tarski(k2_tarski(X23,X24))=k2_xboole_0(X23,X24), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
cnf(c_0_75, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_76, plain, (X1=k1_xboole_0|r2_hidden(esk5_3(k1_xboole_0,X2,X1),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_51])).
cnf(c_0_77, plain, (r1_xboole_0(X1,k4_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_78, negated_conjecture, (r1_tarski(X1,esk2_0)|~r1_tarski(X1,k5_xboole_0(esk1_0,k1_setfam_1(k1_enumset1(esk1_0,esk1_0,X2))))), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_79, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_72])).
cnf(c_0_80, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=X1|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_73, c_0_29])).
fof(c_0_81, plain, ![X8, X9]:r1_tarski(X8,k2_xboole_0(X8,X9)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_82, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_74])).
fof(c_0_83, plain, ![X33, X34, X35]:(~v1_relat_1(X35)|k10_relat_1(X35,k2_xboole_0(X33,X34))=k2_xboole_0(k10_relat_1(X35,X33),k10_relat_1(X35,X34))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t175_relat_1])])).
fof(c_0_84, plain, ![X58, X59]:k2_xboole_0(X58,k4_xboole_0(X59,X58))=k2_xboole_0(X58,X59), inference(variable_rename,[status(thm)],[t39_xboole_1])).
cnf(c_0_85, plain, (X1=k1_xboole_0|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_86, plain, (r1_xboole_0(X1,k5_xboole_0(X3,k1_setfam_1(k1_enumset1(X3,X3,X2))))|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_77, c_0_32])).
cnf(c_0_87, negated_conjecture, (r1_tarski(k5_xboole_0(esk1_0,k1_setfam_1(k1_enumset1(esk1_0,esk1_0,X1))),esk2_0)), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_88, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=X2|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_80, c_0_56])).
fof(c_0_89, plain, ![X57]:k2_xboole_0(X57,k1_xboole_0)=X57, inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_90, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_91, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_82, c_0_26])).
cnf(c_0_92, plain, (k10_relat_1(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k10_relat_1(X1,X2),k10_relat_1(X1,X3))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_93, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_94, plain, (k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X2)))=k1_xboole_0|~r1_tarski(k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X2))),X2)), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_95, negated_conjecture, (r1_tarski(k5_xboole_0(esk1_0,X1),esk2_0)|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_96, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_97, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_98, plain, (k10_relat_1(X1,k3_tarski(k1_enumset1(X2,X2,X3)))=k3_tarski(k1_enumset1(k10_relat_1(X1,X2),k10_relat_1(X1,X2),k10_relat_1(X1,X3)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_91]), c_0_91])).
cnf(c_0_99, plain, (k3_tarski(k1_enumset1(X1,X1,k5_xboole_0(X2,k1_setfam_1(k1_enumset1(X2,X2,X1)))))=k3_tarski(k1_enumset1(X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_91]), c_0_91]), c_0_32])).
cnf(c_0_100, negated_conjecture, (k5_xboole_0(esk1_0,k1_setfam_1(k1_enumset1(esk1_0,esk1_0,esk2_0)))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_43])])).
cnf(c_0_101, plain, (k3_tarski(k1_enumset1(X1,X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[c_0_96, c_0_91])).
cnf(c_0_102, plain, (r1_tarski(k10_relat_1(X1,X2),k10_relat_1(X1,k3_tarski(k1_enumset1(X2,X2,X3))))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_103, negated_conjecture, (k3_tarski(k1_enumset1(esk1_0,esk1_0,esk2_0))=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_101]), c_0_56])).
cnf(c_0_104, negated_conjecture, (~r1_tarski(k10_relat_1(esk3_0,esk1_0),k10_relat_1(esk3_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_105, negated_conjecture, (r1_tarski(k10_relat_1(X1,esk1_0),k10_relat_1(X1,esk2_0))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_106, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_107, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_106])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 108
# Proof object clause steps            : 63
# Proof object formula steps           : 45
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 22
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 30
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 37
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 33
# Processed clauses                    : 1663
# ...of these trivial                  : 20
# ...subsumed                          : 1254
# ...remaining for further processing  : 389
# Other redundant clauses eliminated   : 36
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 21
# Backward-rewritten                   : 18
# Generated clauses                    : 11784
# ...of the previous two non-trivial   : 10915
# Contextual simplify-reflections      : 6
# Paramodulations                      : 11724
# Factorizations                       : 24
# Equation resolutions                 : 36
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 309
#    Positive orientable unit clauses  : 30
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 275
# Current number of unprocessed clauses: 9256
# ...number of literals in the above   : 36614
# Current number of archived formulas  : 0
# Current number of archived clauses   : 75
# Clause-clause subsumption calls (NU) : 35021
# Rec. Clause-clause subsumption calls : 24524
# Non-unit clause-clause subsumptions  : 1035
# Unit Clause-clause subsumption calls : 317
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 92
# BW rewrite match successes           : 29
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 212874

# -------------------------------------------------
# User time                : 0.142 s
# System time              : 0.010 s
# Total time               : 0.151 s
# Maximum resident set size: 28608 pages
