# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t36_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>r1_lattice3(X1,k4_waybel_0(X1,X2),X3))))), file('data/MPTP/proofs/7185c5339d964179a0fc662de4fada26.E_input', t36_waybel_0)).
fof(t15_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k4_waybel_0(X1,X2)=a_2_1_waybel_0(X1,X2))), file('data/MPTP/proofs/7185c5339d964179a0fc662de4fada26.E_input', t15_waybel_0)).
fof(fraenkel_a_2_1_waybel_0, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&l1_orders_2(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r2_hidden(X1,a_2_1_waybel_0(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X5,X4)&r2_hidden(X5,X3))))))), file('data/MPTP/proofs/7185c5339d964179a0fc662de4fada26.E_input', fraenkel_a_2_1_waybel_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/7185c5339d964179a0fc662de4fada26.E_input', t24_orders_2)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/7185c5339d964179a0fc662de4fada26.E_input', d8_lattice3)).
fof(t4_yellow_0, axiom, ![X1]:((v4_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)=>![X4]:((r1_lattice3(X1,X4,X3)=>r1_lattice3(X1,X4,X2))&(r2_lattice3(X1,X4,X2)=>r2_lattice3(X1,X4,X3))))))), file('data/MPTP/proofs/7185c5339d964179a0fc662de4fada26.E_input', t4_yellow_0)).
fof(t7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2)))))))), file('data/MPTP/proofs/7185c5339d964179a0fc662de4fada26.E_input', t7_yellow_0)).
fof(c_0_7, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>r1_lattice3(X1,k4_waybel_0(X1,X2),X3)))))), inference(assume_negation,[status(cth)],[t36_waybel_0])).
fof(c_0_8, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&((~r1_lattice3(esk1_0,esk2_0,esk3_0)|~r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0))&(r1_lattice3(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_7])])])])).
fof(c_0_9, plain, ![X9, X10]:(v2_struct_0(X9)|~l1_orders_2(X9)|(~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|k4_waybel_0(X9,X10)=a_2_1_waybel_0(X9,X10))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t15_waybel_0])])])])).
fof(c_0_10, plain, ![X11, X12, X13, X16, X17]:(((m1_subset_1(esk4_3(X11,X12,X13),u1_struct_0(X12))|~r2_hidden(X11,a_2_1_waybel_0(X12,X13))|(v2_struct_0(X12)|~l1_orders_2(X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))))&((X11=esk4_3(X11,X12,X13)|~r2_hidden(X11,a_2_1_waybel_0(X12,X13))|(v2_struct_0(X12)|~l1_orders_2(X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))))&((m1_subset_1(esk5_3(X11,X12,X13),u1_struct_0(X12))|~r2_hidden(X11,a_2_1_waybel_0(X12,X13))|(v2_struct_0(X12)|~l1_orders_2(X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))))&((r1_orders_2(X12,esk5_3(X11,X12,X13),esk4_3(X11,X12,X13))|~r2_hidden(X11,a_2_1_waybel_0(X12,X13))|(v2_struct_0(X12)|~l1_orders_2(X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))))&(r2_hidden(esk5_3(X11,X12,X13),X13)|~r2_hidden(X11,a_2_1_waybel_0(X12,X13))|(v2_struct_0(X12)|~l1_orders_2(X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))))))))&(~m1_subset_1(X16,u1_struct_0(X12))|(X11!=X16|(~m1_subset_1(X17,u1_struct_0(X12))|(~r1_orders_2(X12,X17,X16)|~r2_hidden(X17,X13))))|r2_hidden(X11,a_2_1_waybel_0(X12,X13))|(v2_struct_0(X12)|~l1_orders_2(X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_1_waybel_0])])])])])])).
fof(c_0_11, plain, ![X18, X19]:(v2_struct_0(X18)|(~v3_orders_2(X18)|~l1_orders_2(X18))|(~m1_subset_1(X19,u1_struct_0(X18))|r1_orders_2(X18,X19,X19))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_12, plain, ![X20, X21, X22, X23]:((~r1_lattice3(X20,X21,X22)|(~m1_subset_1(X23,u1_struct_0(X20))|(~r2_hidden(X23,X21)|r1_orders_2(X20,X22,X23)))|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))&((m1_subset_1(esk6_3(X20,X21,X22),u1_struct_0(X20))|r1_lattice3(X20,X21,X22)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))&((r2_hidden(esk6_3(X20,X21,X22),X21)|r1_lattice3(X20,X21,X22)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))&(~r1_orders_2(X20,X22,esk6_3(X20,X21,X22))|r1_lattice3(X20,X21,X22)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
cnf(c_0_13, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, plain, (v2_struct_0(X1)|k4_waybel_0(X1,X2)=a_2_1_waybel_0(X1,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_17, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_18, plain, (r2_hidden(X3,a_2_1_waybel_0(X2,X5))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~m1_subset_1(X4,u1_struct_0(X2))|~r1_orders_2(X2,X4,X1)|~r2_hidden(X4,X5)|~l1_orders_2(X2)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_21, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22, negated_conjecture, (r1_lattice3(esk1_0,a_2_1_waybel_0(esk1_0,esk2_0),esk3_0)|r1_lattice3(esk1_0,esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16])]), c_0_17])).
cnf(c_0_23, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_24, plain, (r2_hidden(X1,a_2_1_waybel_0(X2,X3))|v2_struct_0(X2)|~r1_orders_2(X2,X4,X1)|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(er,[status(thm)],[c_0_18])).
cnf(c_0_25, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_16])]), c_0_17])).
cnf(c_0_26, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,X1)|r1_lattice3(esk1_0,esk2_0,esk3_0)|~r2_hidden(X1,a_2_1_waybel_0(esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_16])])).
cnf(c_0_27, negated_conjecture, (r2_hidden(X1,a_2_1_waybel_0(esk1_0,X2))|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_16])]), c_0_17])).
cnf(c_0_28, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,X1)|r1_lattice3(esk1_0,esk2_0,esk3_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_15])])).
cnf(c_0_29, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_30, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk6_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_31, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk6_3(X1,esk2_0,X2))|r1_lattice3(esk1_0,esk2_0,esk3_0)|r1_lattice3(X1,esk2_0,X2)|~m1_subset_1(esk6_3(X1,esk2_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_32, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk3_0)|~m1_subset_1(esk6_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_23]), c_0_16])])).
cnf(c_0_33, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_34, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_23]), c_0_16])])).
fof(c_0_35, plain, ![X25, X26, X27, X28]:((~r1_lattice3(X25,X28,X27)|r1_lattice3(X25,X28,X26)|~r1_orders_2(X25,X26,X27)|~m1_subset_1(X27,u1_struct_0(X25))|~m1_subset_1(X26,u1_struct_0(X25))|(~v4_orders_2(X25)|~l1_orders_2(X25)))&(~r2_lattice3(X25,X28,X26)|r2_lattice3(X25,X28,X27)|~r1_orders_2(X25,X26,X27)|~m1_subset_1(X27,u1_struct_0(X25))|~m1_subset_1(X26,u1_struct_0(X25))|(~v4_orders_2(X25)|~l1_orders_2(X25)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_yellow_0])])])])).
cnf(c_0_36, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,X1)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_34]), c_0_23]), c_0_16])])).
cnf(c_0_37, plain, (r2_hidden(esk5_3(X1,X2,X3),X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_38, plain, (r1_lattice3(X1,X2,X4)|~r1_lattice3(X1,X2,X3)|~r1_orders_2(X1,X4,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_39, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk5_3(X1,X2,esk2_0))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,esk2_0))|~m1_subset_1(esk5_3(X1,X2,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_40, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_41, plain, (r1_orders_2(X1,esk5_3(X2,X1,X3),esk4_3(X2,X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,a_2_1_waybel_0(X1,X3))|~l1_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_42, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_43, plain, (m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_44, negated_conjecture, (r1_lattice3(esk1_0,X1,esk3_0)|v2_struct_0(X2)|~r2_hidden(X3,a_2_1_waybel_0(X2,esk2_0))|~r1_lattice3(esk1_0,X1,esk5_3(X3,X2,esk2_0))|~m1_subset_1(esk5_3(X3,X2,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_23]), c_0_16]), c_0_40])])).
cnf(c_0_45, plain, (r1_lattice3(X1,X2,esk5_3(X3,X1,X4))|v2_struct_0(X1)|~r2_hidden(X3,a_2_1_waybel_0(X1,X4))|~r1_lattice3(X1,X2,esk4_3(X3,X1,X4))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v4_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_41]), c_0_42]), c_0_43])).
fof(c_0_46, plain, ![X29, X30, X31]:((~r1_lattice3(X29,k1_tarski(X31),X30)|r1_orders_2(X29,X30,X31)|~m1_subset_1(X31,u1_struct_0(X29))|~m1_subset_1(X30,u1_struct_0(X29))|~l1_orders_2(X29))&((~r1_orders_2(X29,X30,X31)|r1_lattice3(X29,k1_tarski(X31),X30)|~m1_subset_1(X31,u1_struct_0(X29))|~m1_subset_1(X30,u1_struct_0(X29))|~l1_orders_2(X29))&((~r2_lattice3(X29,k1_tarski(X31),X30)|r1_orders_2(X29,X31,X30)|~m1_subset_1(X31,u1_struct_0(X29))|~m1_subset_1(X30,u1_struct_0(X29))|~l1_orders_2(X29))&(~r1_orders_2(X29,X31,X30)|r2_lattice3(X29,k1_tarski(X31),X30)|~m1_subset_1(X31,u1_struct_0(X29))|~m1_subset_1(X30,u1_struct_0(X29))|~l1_orders_2(X29))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])).
cnf(c_0_47, negated_conjecture, (r1_lattice3(esk1_0,X1,esk3_0)|~r2_hidden(X2,a_2_1_waybel_0(esk1_0,esk2_0))|~r1_lattice3(esk1_0,X1,esk4_3(X2,esk1_0,esk2_0))|~m1_subset_1(esk5_3(X2,esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_15]), c_0_16]), c_0_40])]), c_0_17])).
cnf(c_0_48, plain, (r1_lattice3(X1,k1_tarski(X3),X2)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_49, negated_conjecture, (r1_lattice3(esk1_0,X1,esk3_0)|~r2_hidden(X2,a_2_1_waybel_0(esk1_0,esk2_0))|~r1_lattice3(esk1_0,X1,esk4_3(X2,esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_43]), c_0_15]), c_0_16])]), c_0_17])).
cnf(c_0_50, negated_conjecture, (r1_lattice3(esk1_0,k1_tarski(X1),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_25]), c_0_16])])).
cnf(c_0_51, negated_conjecture, (r1_lattice3(esk1_0,k1_tarski(esk4_3(X1,esk1_0,esk2_0)),esk3_0)|~r2_hidden(X1,a_2_1_waybel_0(esk1_0,esk2_0))|~m1_subset_1(esk4_3(X1,esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_52, plain, (X1=esk4_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_53, plain, (r1_orders_2(X1,X3,X2)|~r1_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_54, negated_conjecture, (r1_lattice3(esk1_0,k1_tarski(X1),esk3_0)|~r2_hidden(X1,a_2_1_waybel_0(esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_15]), c_0_16])]), c_0_17])).
cnf(c_0_55, negated_conjecture, (~r1_lattice3(esk1_0,esk2_0,esk3_0)|~r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_56, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,X1)|~r2_hidden(X1,a_2_1_waybel_0(esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_23]), c_0_16])])).
cnf(c_0_57, negated_conjecture, (~r1_lattice3(esk1_0,a_2_1_waybel_0(esk1_0,esk2_0),esk3_0)|~r1_lattice3(esk1_0,esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_14]), c_0_15]), c_0_16])]), c_0_17])).
cnf(c_0_58, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk6_3(X1,a_2_1_waybel_0(esk1_0,esk2_0),X2))|r1_lattice3(X1,a_2_1_waybel_0(esk1_0,esk2_0),X2)|~m1_subset_1(esk6_3(X1,a_2_1_waybel_0(esk1_0,esk2_0),X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_56, c_0_29])).
cnf(c_0_59, negated_conjecture, (~r1_lattice3(esk1_0,a_2_1_waybel_0(esk1_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_34])])).
cnf(c_0_60, negated_conjecture, (~m1_subset_1(esk6_3(esk1_0,a_2_1_waybel_0(esk1_0,esk2_0),esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_58]), c_0_23]), c_0_16])]), c_0_59])).
cnf(c_0_61, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_33]), c_0_23]), c_0_16])]), c_0_59]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 62
# Proof object clause steps            : 47
# Proof object formula steps           : 15
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 22
# Proof object simplifying inferences  : 63
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 26
# Processed clauses                    : 338
# ...of these trivial                  : 0
# ...subsumed                          : 111
# ...remaining for further processing  : 227
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 18
# Backward-rewritten                   : 10
# Generated clauses                    : 626
# ...of the previous two non-trivial   : 607
# Contextual simplify-reflections      : 14
# Paramodulations                      : 625
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 172
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 162
# Current number of unprocessed clauses: 305
# ...number of literals in the above   : 2721
# Current number of archived formulas  : 0
# Current number of archived clauses   : 54
# Clause-clause subsumption calls (NU) : 10661
# Rec. Clause-clause subsumption calls : 901
# Non-unit clause-clause subsumptions  : 142
# Unit Clause-clause subsumption calls : 51
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 29087

# -------------------------------------------------
# User time                : 0.027 s
# System time              : 0.003 s
# Total time               : 0.030 s
# Maximum resident set size: 28624 pages
