# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t44_tex_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v2_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v2_tex_2(X2,X1)<=>v1_zfmisc_1(X2)))), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', t44_tex_2)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', t3_subset)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', d3_tarski)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', d1_xboole_0)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', t1_subset)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', cc1_subset_1)).
fof(t49_zfmisc_1, axiom, ![X1, X2]:~(k2_xboole_0(k1_tarski(X1),X2)=k1_xboole_0), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', t49_zfmisc_1)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', t1_boole)).
fof(t36_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>v2_tex_2(k6_domain_1(u1_struct_0(X1),X2),X1))), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', t36_tex_2)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', redefinition_k6_domain_1)).
fof(l1_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', l1_zfmisc_1)).
fof(t8_boole, axiom, ![X1, X2]:~((v1_xboole_0(X1)&(~(X1=X2)&v1_xboole_0(X2)))), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', t8_boole)).
fof(t1_tex_2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((~(v1_xboole_0(X2))&v1_zfmisc_1(X2))=>(r1_tarski(X1,X2)=>X1=X2))), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', t1_tex_2)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', fc1_xboole_0)).
fof(t34_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>~((v2_tex_2(X2,X1)&![X3]:((~(v2_struct_0(X3))&(v1_pre_topc(X3)&(v1_tdlat_3(X3)&m1_pre_topc(X3,X1))))=>~(X2=u1_struct_0(X3))))))), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', t34_tex_2)).
fof(fc7_struct_0, axiom, ![X1]:((v7_struct_0(X1)&l1_struct_0(X1))=>v1_zfmisc_1(u1_struct_0(X1))), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', fc7_struct_0)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', dt_m1_pre_topc)).
fof(cc6_tdlat_3, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v2_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:(m1_pre_topc(X2,X1)=>v2_tdlat_3(X2))), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', cc6_tdlat_3)).
fof(cc2_tdlat_3, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_tdlat_3(X1)=>v2_pre_topc(X1))), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', cc2_tdlat_3)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', dt_l1_pre_topc)).
fof(cc2_tex_1, axiom, ![X1]:(l1_pre_topc(X1)=>((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v1_tdlat_3(X1)&v2_tdlat_3(X1))))=>(~(v2_struct_0(X1))&(v7_struct_0(X1)&v2_pre_topc(X1))))), file('data/MPTP/proofs/71a0782dd75f41c0a67cd3815e931c80.E_input', cc2_tex_1)).
fof(c_0_21, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v2_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v2_tex_2(X2,X1)<=>v1_zfmisc_1(X2))))), inference(assume_negation,[status(cth)],[t44_tex_2])).
fof(c_0_22, plain, ![X15, X16]:((~m1_subset_1(X15,k1_zfmisc_1(X16))|r1_tarski(X15,X16))&(~r1_tarski(X15,X16)|m1_subset_1(X15,k1_zfmisc_1(X16)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_23, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&(v2_tdlat_3(esk1_0)&l1_pre_topc(esk1_0))))&((~v1_xboole_0(esk2_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))&((~v2_tex_2(esk2_0,esk1_0)|~v1_zfmisc_1(esk2_0))&(v2_tex_2(esk2_0,esk1_0)|v1_zfmisc_1(esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_21])])])])).
fof(c_0_24, plain, ![X22, X23, X24, X25, X26]:((~r1_tarski(X22,X23)|(~r2_hidden(X24,X22)|r2_hidden(X24,X23)))&((r2_hidden(esk4_2(X25,X26),X25)|r1_tarski(X25,X26))&(~r2_hidden(esk4_2(X25,X26),X26)|r1_tarski(X25,X26)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_25, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_28, negated_conjecture, (r1_tarski(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_29, plain, ![X8, X9, X10]:((~v1_xboole_0(X8)|~r2_hidden(X9,X8))&(r2_hidden(esk3_1(X10),X10)|v1_xboole_0(X10))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_30, plain, ![X12, X13]:(~r2_hidden(X12,X13)|m1_subset_1(X12,X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_31, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_32, plain, (r2_hidden(esk3_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_33, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_34, plain, ![X43, X44]:(~v1_xboole_0(X43)|(~m1_subset_1(X44,k1_zfmisc_1(X43))|v1_xboole_0(X44))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
fof(c_0_35, plain, ![X18, X19]:k2_xboole_0(k1_tarski(X18),X19)!=k1_xboole_0, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t49_zfmisc_1])])).
fof(c_0_36, plain, ![X17]:k2_xboole_0(X17,k1_xboole_0)=X17, inference(variable_rename,[status(thm)],[t1_boole])).
fof(c_0_37, plain, ![X35, X36]:(v2_struct_0(X35)|(~v2_pre_topc(X35)|~l1_pre_topc(X35))|(~m1_subset_1(X36,u1_struct_0(X35))|v2_tex_2(k6_domain_1(u1_struct_0(X35),X36),X35))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t36_tex_2])])])])).
cnf(c_0_38, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39, negated_conjecture, (r2_hidden(esk3_1(esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])).
fof(c_0_40, plain, ![X33, X34]:(v1_xboole_0(X33)|~m1_subset_1(X34,X33)|k6_domain_1(X33,X34)=k1_tarski(X34)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_41, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_42, plain, ![X38, X39]:((~r1_tarski(k1_tarski(X38),X39)|r2_hidden(X38,X39))&(~r2_hidden(X38,X39)|r1_tarski(k1_tarski(X38),X39))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])).
cnf(c_0_43, plain, (k2_xboole_0(k1_tarski(X1),X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_44, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_45, plain, ![X45, X46]:(~v1_xboole_0(X45)|(X45=X46|~v1_xboole_0(X46))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_boole])])])).
cnf(c_0_46, plain, (v2_struct_0(X1)|v2_tex_2(k6_domain_1(u1_struct_0(X1),X2),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_47, negated_conjecture, (m1_subset_1(esk3_1(esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_48, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_49, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_50, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_51, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_52, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_26]), c_0_33])).
fof(c_0_53, plain, ![X29, X30]:(v1_xboole_0(X29)|(v1_xboole_0(X30)|~v1_zfmisc_1(X30)|(~r1_tarski(X29,X30)|X29=X30))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t1_tex_2])])])])).
cnf(c_0_54, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_55, plain, (k1_tarski(X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_56, plain, (X1=X2|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_57, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
fof(c_0_58, plain, ![X40, X41]:(((~v2_struct_0(esk5_2(X40,X41))|~v2_tex_2(X41,X40)|(v1_xboole_0(X41)|~m1_subset_1(X41,k1_zfmisc_1(u1_struct_0(X40))))|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40))))&((v1_pre_topc(esk5_2(X40,X41))|~v2_tex_2(X41,X40)|(v1_xboole_0(X41)|~m1_subset_1(X41,k1_zfmisc_1(u1_struct_0(X40))))|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40))))&((v1_tdlat_3(esk5_2(X40,X41))|~v2_tex_2(X41,X40)|(v1_xboole_0(X41)|~m1_subset_1(X41,k1_zfmisc_1(u1_struct_0(X40))))|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40))))&(m1_pre_topc(esk5_2(X40,X41),X40)|~v2_tex_2(X41,X40)|(v1_xboole_0(X41)|~m1_subset_1(X41,k1_zfmisc_1(u1_struct_0(X40))))|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40)))))))&(X41=u1_struct_0(esk5_2(X40,X41))|~v2_tex_2(X41,X40)|(v1_xboole_0(X41)|~m1_subset_1(X41,k1_zfmisc_1(u1_struct_0(X40))))|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_tex_2])])])])])])).
cnf(c_0_59, negated_conjecture, (v2_tex_2(k6_domain_1(u1_struct_0(esk1_0),esk3_1(esk2_0)),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48]), c_0_49])]), c_0_50])).
cnf(c_0_60, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk3_1(esk2_0))=k1_tarski(esk3_1(esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_47]), c_0_52])).
cnf(c_0_61, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|X1=X2|~v1_zfmisc_1(X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_62, plain, (r1_tarski(k1_tarski(esk3_1(X1)),X1)|v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_54, c_0_32])).
cnf(c_0_63, plain, (~v1_xboole_0(k1_tarski(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56])]), c_0_57])])).
cnf(c_0_64, plain, (X1=u1_struct_0(esk5_2(X2,X1))|v1_xboole_0(X1)|v2_struct_0(X2)|~v2_tex_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_65, negated_conjecture, (v2_tex_2(k1_tarski(esk3_1(esk2_0)),esk1_0)), inference(rw,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_66, plain, (k1_tarski(esk3_1(X1))=X1|v1_xboole_0(X1)|~v1_zfmisc_1(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63])).
cnf(c_0_67, negated_conjecture, (v2_tex_2(esk2_0,esk1_0)|v1_zfmisc_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_68, plain, (m1_pre_topc(esk5_2(X1,X2),X1)|v1_xboole_0(X2)|v2_struct_0(X1)|~v2_tex_2(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
fof(c_0_69, plain, ![X14]:(~v7_struct_0(X14)|~l1_struct_0(X14)|v1_zfmisc_1(u1_struct_0(X14))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc7_struct_0])])).
cnf(c_0_70, negated_conjecture, (u1_struct_0(esk5_2(esk1_0,esk2_0))=esk2_0|~v2_tex_2(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_26]), c_0_48]), c_0_49])]), c_0_33]), c_0_50])).
cnf(c_0_71, negated_conjecture, (v2_tex_2(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_33]), c_0_67])).
cnf(c_0_72, negated_conjecture, (~v2_tex_2(esk2_0,esk1_0)|~v1_zfmisc_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_73, plain, ![X31, X32]:(~l1_pre_topc(X31)|(~m1_pre_topc(X32,X31)|l1_pre_topc(X32))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_74, negated_conjecture, (m1_pre_topc(esk5_2(esk1_0,esk2_0),esk1_0)|~v2_tex_2(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_26]), c_0_48]), c_0_49])]), c_0_33]), c_0_50])).
fof(c_0_75, plain, ![X6, X7]:(v2_struct_0(X6)|(~v2_pre_topc(X6)|(~v2_tdlat_3(X6)|~l1_pre_topc(X6)))|(~m1_pre_topc(X7,X6)|v2_tdlat_3(X7))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc6_tdlat_3])])])])).
fof(c_0_76, plain, ![X21]:(~l1_pre_topc(X21)|(~v2_tdlat_3(X21)|v2_pre_topc(X21))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_tdlat_3])])).
cnf(c_0_77, plain, (v1_zfmisc_1(u1_struct_0(X1))|~v7_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_78, negated_conjecture, (u1_struct_0(esk5_2(esk1_0,esk2_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_71])])).
cnf(c_0_79, negated_conjecture, (~v1_zfmisc_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_71])])).
fof(c_0_80, plain, ![X28]:(~l1_pre_topc(X28)|l1_struct_0(X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_81, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_82, negated_conjecture, (m1_pre_topc(esk5_2(esk1_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_71])])).
fof(c_0_83, plain, ![X20]:((~v2_struct_0(X20)|(v2_struct_0(X20)|(~v2_pre_topc(X20)|(~v1_tdlat_3(X20)|~v2_tdlat_3(X20))))|~l1_pre_topc(X20))&((v7_struct_0(X20)|(v2_struct_0(X20)|(~v2_pre_topc(X20)|(~v1_tdlat_3(X20)|~v2_tdlat_3(X20))))|~l1_pre_topc(X20))&(v2_pre_topc(X20)|(v2_struct_0(X20)|(~v2_pre_topc(X20)|(~v1_tdlat_3(X20)|~v2_tdlat_3(X20))))|~l1_pre_topc(X20)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_tex_1])])])])).
cnf(c_0_84, plain, (v1_tdlat_3(esk5_2(X1,X2))|v1_xboole_0(X2)|v2_struct_0(X1)|~v2_tex_2(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_85, plain, (v2_struct_0(X1)|v2_tdlat_3(X2)|~v2_pre_topc(X1)|~v2_tdlat_3(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_86, plain, (v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_tdlat_3(X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_87, negated_conjecture, (~l1_struct_0(esk5_2(esk1_0,esk2_0))|~v7_struct_0(esk5_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_79])).
cnf(c_0_88, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_89, negated_conjecture, (l1_pre_topc(esk5_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_48])])).
cnf(c_0_90, plain, (v7_struct_0(X1)|v2_struct_0(X1)|~v2_pre_topc(X1)|~v1_tdlat_3(X1)|~v2_tdlat_3(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_91, negated_conjecture, (v1_tdlat_3(esk5_2(esk1_0,esk2_0))|~v2_tex_2(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_26]), c_0_48]), c_0_49])]), c_0_33]), c_0_50])).
cnf(c_0_92, plain, (v2_tdlat_3(X1)|v2_struct_0(X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_tdlat_3(X2)), inference(csr,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_93, negated_conjecture, (v2_tdlat_3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_94, negated_conjecture, (~v7_struct_0(esk5_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_89])])).
cnf(c_0_95, plain, (v7_struct_0(X1)|v2_struct_0(X1)|~v1_tdlat_3(X1)|~l1_pre_topc(X1)|~v2_tdlat_3(X1)), inference(csr,[status(thm)],[c_0_90, c_0_86])).
cnf(c_0_96, negated_conjecture, (v1_tdlat_3(esk5_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_71])])).
cnf(c_0_97, negated_conjecture, (v2_tdlat_3(esk5_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_82]), c_0_48]), c_0_93])]), c_0_50])).
cnf(c_0_98, plain, (v1_xboole_0(X2)|v2_struct_0(X1)|~v2_struct_0(esk5_2(X1,X2))|~v2_tex_2(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_99, negated_conjecture, (v2_struct_0(esk5_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_89])]), c_0_96]), c_0_97])])).
cnf(c_0_100, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_71]), c_0_26]), c_0_48]), c_0_49])]), c_0_33]), c_0_50]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 101
# Proof object clause steps            : 59
# Proof object formula steps           : 42
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 21
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 60
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 38
# Processed clauses                    : 310
# ...of these trivial                  : 0
# ...subsumed                          : 141
# ...remaining for further processing  : 169
# Other redundant clauses eliminated   : 11
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 7
# Generated clauses                    : 421
# ...of the previous two non-trivial   : 366
# Contextual simplify-reflections      : 3
# Paramodulations                      : 410
# Factorizations                       : 0
# Equation resolutions                 : 11
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 121
#    Positive orientable unit clauses  : 30
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 82
# Current number of unprocessed clauses: 126
# ...number of literals in the above   : 614
# Current number of archived formulas  : 0
# Current number of archived clauses   : 48
# Clause-clause subsumption calls (NU) : 2552
# Rec. Clause-clause subsumption calls : 1030
# Non-unit clause-clause subsumptions  : 78
# Unit Clause-clause subsumption calls : 162
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 11
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7844

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 28820 pages
