# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t134_enumset1, conjecture, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8),k1_tarski(X9)), file('data/MPTP/proofs/71c84d58ec08476aac3befa0a7992e40.E_input', t134_enumset1)).
fof(t127_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k1_tarski(X1),k6_enumset1(X2,X3,X4,X5,X6,X7,X8,X9)), file('data/MPTP/proofs/71c84d58ec08476aac3befa0a7992e40.E_input', t127_enumset1)).
fof(t129_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k1_enumset1(X1,X2,X3),k4_enumset1(X4,X5,X6,X7,X8,X9)), file('data/MPTP/proofs/71c84d58ec08476aac3befa0a7992e40.E_input', t129_enumset1)).
fof(t132_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k4_enumset1(X1,X2,X3,X4,X5,X6),k1_enumset1(X7,X8,X9)), file('data/MPTP/proofs/71c84d58ec08476aac3befa0a7992e40.E_input', t132_enumset1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/71c84d58ec08476aac3befa0a7992e40.E_input', commutativity_k2_xboole_0)).
fof(t128_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k2_tarski(X1,X2),k5_enumset1(X3,X4,X5,X6,X7,X8,X9)), file('data/MPTP/proofs/71c84d58ec08476aac3befa0a7992e40.E_input', t128_enumset1)).
fof(t133_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k5_enumset1(X1,X2,X3,X4,X5,X6,X7),k2_tarski(X8,X9)), file('data/MPTP/proofs/71c84d58ec08476aac3befa0a7992e40.E_input', t133_enumset1)).
fof(c_0_7, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6, X7, X8, X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8),k1_tarski(X9))), inference(assume_negation,[status(cth)],[t134_enumset1])).
fof(c_0_8, negated_conjecture, k7_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0,esk9_0)!=k2_xboole_0(k6_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0),k1_tarski(esk9_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
fof(c_0_9, plain, ![X19, X20, X21, X22, X23, X24, X25, X26, X27]:k7_enumset1(X19,X20,X21,X22,X23,X24,X25,X26,X27)=k2_xboole_0(k1_tarski(X19),k6_enumset1(X20,X21,X22,X23,X24,X25,X26,X27)), inference(variable_rename,[status(thm)],[t127_enumset1])).
fof(c_0_10, plain, ![X28, X29, X30, X31, X32, X33, X34, X35, X36]:k7_enumset1(X28,X29,X30,X31,X32,X33,X34,X35,X36)=k2_xboole_0(k1_enumset1(X28,X29,X30),k4_enumset1(X31,X32,X33,X34,X35,X36)), inference(variable_rename,[status(thm)],[t129_enumset1])).
fof(c_0_11, plain, ![X37, X38, X39, X40, X41, X42, X43, X44, X45]:k7_enumset1(X37,X38,X39,X40,X41,X42,X43,X44,X45)=k2_xboole_0(k4_enumset1(X37,X38,X39,X40,X41,X42),k1_enumset1(X43,X44,X45)), inference(variable_rename,[status(thm)],[t132_enumset1])).
cnf(c_0_12, negated_conjecture, (k7_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0,esk9_0)!=k2_xboole_0(k6_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0),k1_tarski(esk9_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k1_tarski(X1),k6_enumset1(X2,X3,X4,X5,X6,X7,X8,X9))), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_14, plain, ![X46, X47]:k2_xboole_0(X46,X47)=k2_xboole_0(X47,X46), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_15, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k1_enumset1(X1,X2,X3),k4_enumset1(X4,X5,X6,X7,X8,X9))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k4_enumset1(X1,X2,X3,X4,X5,X6),k1_enumset1(X7,X8,X9))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_17, plain, ![X48, X49, X50, X51, X52, X53, X54, X55, X56]:k7_enumset1(X48,X49,X50,X51,X52,X53,X54,X55,X56)=k2_xboole_0(k2_tarski(X48,X49),k5_enumset1(X50,X51,X52,X53,X54,X55,X56)), inference(variable_rename,[status(thm)],[t128_enumset1])).
fof(c_0_18, plain, ![X57, X58, X59, X60, X61, X62, X63, X64, X65]:k7_enumset1(X57,X58,X59,X60,X61,X62,X63,X64,X65)=k2_xboole_0(k5_enumset1(X57,X58,X59,X60,X61,X62,X63),k2_tarski(X64,X65)), inference(variable_rename,[status(thm)],[t133_enumset1])).
cnf(c_0_19, negated_conjecture, (k2_xboole_0(k1_tarski(esk1_0),k6_enumset1(esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0,esk9_0))!=k2_xboole_0(k6_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0),k1_tarski(esk9_0))), inference(rw,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_20, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (k2_xboole_0(k1_tarski(X1),k6_enumset1(X2,X3,X4,X5,X6,X7,X8,X9))=k2_xboole_0(k1_enumset1(X1,X2,X3),k4_enumset1(X4,X5,X6,X7,X8,X9))), inference(rw,[status(thm)],[c_0_15, c_0_13])).
cnf(c_0_22, plain, (k2_xboole_0(k1_tarski(X1),k6_enumset1(X2,X3,X4,X5,X6,X7,X8,X9))=k2_xboole_0(k4_enumset1(X1,X2,X3,X4,X5,X6),k1_enumset1(X7,X8,X9))), inference(rw,[status(thm)],[c_0_16, c_0_13])).
cnf(c_0_23, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k2_tarski(X1,X2),k5_enumset1(X3,X4,X5,X6,X7,X8,X9))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k5_enumset1(X1,X2,X3,X4,X5,X6,X7),k2_tarski(X8,X9))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, negated_conjecture, (k2_xboole_0(k1_tarski(esk9_0),k6_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0))!=k2_xboole_0(k1_tarski(esk1_0),k6_enumset1(esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0,esk9_0))), inference(rw,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_26, plain, (k2_xboole_0(k1_tarski(X1),k6_enumset1(X2,X3,X4,X5,X6,X7,X8,X9))=k2_xboole_0(k1_tarski(X7),k6_enumset1(X8,X9,X1,X2,X3,X4,X5,X6))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])).
cnf(c_0_27, plain, (k2_xboole_0(k1_tarski(X1),k6_enumset1(X2,X3,X4,X5,X6,X7,X8,X9))=k2_xboole_0(k2_tarski(X1,X2),k5_enumset1(X3,X4,X5,X6,X7,X8,X9))), inference(rw,[status(thm)],[c_0_23, c_0_13])).
cnf(c_0_28, plain, (k2_xboole_0(k1_tarski(X1),k6_enumset1(X2,X3,X4,X5,X6,X7,X8,X9))=k2_xboole_0(k5_enumset1(X1,X2,X3,X4,X5,X6,X7),k2_tarski(X8,X9))), inference(rw,[status(thm)],[c_0_24, c_0_13])).
cnf(c_0_29, negated_conjecture, (k2_xboole_0(k1_tarski(esk3_0),k6_enumset1(esk4_0,esk5_0,esk6_0,esk7_0,esk8_0,esk9_0,esk1_0,esk2_0))!=k2_xboole_0(k1_tarski(esk1_0),k6_enumset1(esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0,esk8_0,esk9_0))), inference(rw,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_30, plain, (k2_xboole_0(k1_tarski(X1),k6_enumset1(X2,X3,X4,X5,X6,X7,X8,X9))=k2_xboole_0(k1_tarski(X8),k6_enumset1(X9,X1,X2,X3,X4,X5,X6,X7))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_27]), c_0_28])).
cnf(c_0_31, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_30])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 32
# Proof object clause steps            : 17
# Proof object formula steps           : 15
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 7
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 2
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 7
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 6
# Processed clauses                    : 19
# ...of these trivial                  : 1
# ...subsumed                          : 2
# ...remaining for further processing  : 16
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 26
# ...of the previous two non-trivial   : 17
# Contextual simplify-reflections      : 0
# Paramodulations                      : 26
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 8
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 10
# ...number of literals in the above   : 10
# Current number of archived formulas  : 0
# Current number of archived clauses   : 9
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 16
# BW rewrite match successes           : 16
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 924

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.000 s
# Total time               : 0.008 s
# Maximum resident set size: 28600 pages
