# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d8_filter_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v3_filter_0(X1)&l3_lattices(X1))))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k4_filter_0(X1,X2,X3)<=>(r3_lattices(X1,k4_lattices(X1,X2,X4),X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r3_lattices(X1,k4_lattices(X1,X2,X5),X3)=>r3_lattices(X1,X5,X4)))))))))), file('data/MPTP/proofs/72327296763e4b3188a433453ce9d0f4.E_input', d8_filter_0)).
fof(dt_k4_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&l1_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k4_lattices(X1,X2,X3),u1_struct_0(X1))), file('data/MPTP/proofs/72327296763e4b3188a433453ce9d0f4.E_input', dt_k4_lattices)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/72327296763e4b3188a433453ce9d0f4.E_input', cc1_lattices)).
fof(dt_l3_lattices, axiom, ![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1))), file('data/MPTP/proofs/72327296763e4b3188a433453ce9d0f4.E_input', dt_l3_lattices)).
fof(dt_k3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v4_lattices(X1)&l2_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k3_lattices(X1,X2,X3),u1_struct_0(X1))), file('data/MPTP/proofs/72327296763e4b3188a433453ce9d0f4.E_input', dt_k3_lattices)).
fof(commutativity_k4_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&l1_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k4_lattices(X1,X2,X3)=k4_lattices(X1,X3,X2)), file('data/MPTP/proofs/72327296763e4b3188a433453ce9d0f4.E_input', commutativity_k4_lattices)).
fof(t3_filter_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_lattices(X1,X2,X3)=>r3_lattices(X1,X2,k3_lattices(X1,X4,X3))))))), file('data/MPTP/proofs/72327296763e4b3188a433453ce9d0f4.E_input', t3_filter_0)).
fof(t56_filter_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v3_filter_0(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_lattices(X1,k4_filter_0(X1,X2,X3),k4_filter_0(X1,X2,k3_lattices(X1,X3,X4)))&(r3_lattices(X1,k4_filter_0(X1,X2,X3),k4_filter_0(X1,k4_lattices(X1,X2,X4),X3))&(r3_lattices(X1,k4_filter_0(X1,X2,X3),k4_filter_0(X1,X2,k3_lattices(X1,X4,X3)))&r3_lattices(X1,k4_filter_0(X1,X2,X3),k4_filter_0(X1,k4_lattices(X1,X4,X2),X3))))))))), file('data/MPTP/proofs/72327296763e4b3188a433453ce9d0f4.E_input', t56_filter_1)).
fof(t2_filter_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_lattices(X1,X2,X3)=>r3_lattices(X1,k4_lattices(X1,X2,X4),X3)))))), file('data/MPTP/proofs/72327296763e4b3188a433453ce9d0f4.E_input', t2_filter_0)).
fof(commutativity_k3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v4_lattices(X1)&l2_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k3_lattices(X1,X2,X3)=k3_lattices(X1,X3,X2)), file('data/MPTP/proofs/72327296763e4b3188a433453ce9d0f4.E_input', commutativity_k3_lattices)).
fof(dt_k4_filter_0, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k4_filter_0(X1,X2,X3),u1_struct_0(X1))), file('data/MPTP/proofs/72327296763e4b3188a433453ce9d0f4.E_input', dt_k4_filter_0)).
fof(t17_filter_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v3_filter_0(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>k4_filter_0(X1,k4_lattices(X1,X2,X3),X4)=k4_filter_0(X1,X2,k4_filter_0(X1,X3,X4)))))), file('data/MPTP/proofs/72327296763e4b3188a433453ce9d0f4.E_input', t17_filter_1)).
fof(c_0_12, plain, ![X6, X7, X8, X9, X10]:(((r3_lattices(X6,k4_lattices(X6,X7,X9),X8)|X9!=k4_filter_0(X6,X7,X8)|~m1_subset_1(X9,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v3_filter_0(X6)|~l3_lattices(X6))))|~m1_subset_1(X8,u1_struct_0(X6))|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|~l3_lattices(X6))))&(~m1_subset_1(X10,u1_struct_0(X6))|(~r3_lattices(X6,k4_lattices(X6,X7,X10),X8)|r3_lattices(X6,X10,X9))|X9!=k4_filter_0(X6,X7,X8)|~m1_subset_1(X9,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v3_filter_0(X6)|~l3_lattices(X6))))|~m1_subset_1(X8,u1_struct_0(X6))|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|~l3_lattices(X6)))))&((m1_subset_1(esk1_4(X6,X7,X8,X9),u1_struct_0(X6))|~r3_lattices(X6,k4_lattices(X6,X7,X9),X8)|X9=k4_filter_0(X6,X7,X8)|~m1_subset_1(X9,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v3_filter_0(X6)|~l3_lattices(X6))))|~m1_subset_1(X8,u1_struct_0(X6))|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|~l3_lattices(X6))))&((r3_lattices(X6,k4_lattices(X6,X7,esk1_4(X6,X7,X8,X9)),X8)|~r3_lattices(X6,k4_lattices(X6,X7,X9),X8)|X9=k4_filter_0(X6,X7,X8)|~m1_subset_1(X9,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v3_filter_0(X6)|~l3_lattices(X6))))|~m1_subset_1(X8,u1_struct_0(X6))|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|~l3_lattices(X6))))&(~r3_lattices(X6,esk1_4(X6,X7,X8,X9),X9)|~r3_lattices(X6,k4_lattices(X6,X7,X9),X8)|X9=k4_filter_0(X6,X7,X8)|~m1_subset_1(X9,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v3_filter_0(X6)|~l3_lattices(X6))))|~m1_subset_1(X8,u1_struct_0(X6))|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|~l3_lattices(X6))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_filter_0])])])])])])).
fof(c_0_13, plain, ![X12, X13, X14]:(v2_struct_0(X12)|(~v6_lattices(X12)|~l1_lattices(X12))|(~m1_subset_1(X13,u1_struct_0(X12))|~m1_subset_1(X14,u1_struct_0(X12)))|m1_subset_1(k4_lattices(X12,X13,X14),u1_struct_0(X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_lattices])])])).
fof(c_0_14, plain, ![X15]:((~v2_struct_0(X15)|(v2_struct_0(X15)|~v10_lattices(X15))|~l3_lattices(X15))&((v4_lattices(X15)|(v2_struct_0(X15)|~v10_lattices(X15))|~l3_lattices(X15))&((v5_lattices(X15)|(v2_struct_0(X15)|~v10_lattices(X15))|~l3_lattices(X15))&((v6_lattices(X15)|(v2_struct_0(X15)|~v10_lattices(X15))|~l3_lattices(X15))&((v7_lattices(X15)|(v2_struct_0(X15)|~v10_lattices(X15))|~l3_lattices(X15))&((v8_lattices(X15)|(v2_struct_0(X15)|~v10_lattices(X15))|~l3_lattices(X15))&(v9_lattices(X15)|(v2_struct_0(X15)|~v10_lattices(X15))|~l3_lattices(X15)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
fof(c_0_15, plain, ![X16]:((l1_lattices(X16)|~l3_lattices(X16))&(l2_lattices(X16)|~l3_lattices(X16))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
fof(c_0_16, plain, ![X17, X18, X19]:(v2_struct_0(X17)|(~v4_lattices(X17)|~l2_lattices(X17))|(~m1_subset_1(X18,u1_struct_0(X17))|~m1_subset_1(X19,u1_struct_0(X17)))|m1_subset_1(k3_lattices(X17,X18,X19),u1_struct_0(X17))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_lattices])])])).
fof(c_0_17, plain, ![X20, X21, X22]:(v2_struct_0(X20)|(~v6_lattices(X20)|~l1_lattices(X20))|(~m1_subset_1(X21,u1_struct_0(X20))|~m1_subset_1(X22,u1_struct_0(X20)))|k4_lattices(X20,X21,X22)=k4_lattices(X20,X22,X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[commutativity_k4_lattices])])])).
cnf(c_0_18, plain, (r3_lattices(X2,X1,X5)|v2_struct_0(X2)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r3_lattices(X2,k4_lattices(X2,X3,X1),X4)|X5!=k4_filter_0(X2,X3,X4)|~m1_subset_1(X5,u1_struct_0(X2))|~v10_lattices(X2)|~v3_filter_0(X2)|~l3_lattices(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~v10_lattices(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_19, plain, ![X23, X24, X25, X26]:(v2_struct_0(X23)|(~v10_lattices(X23)|~l3_lattices(X23))|(~m1_subset_1(X24,u1_struct_0(X23))|(~m1_subset_1(X25,u1_struct_0(X23))|(~m1_subset_1(X26,u1_struct_0(X23))|(~r3_lattices(X23,X24,X25)|r3_lattices(X23,X24,k3_lattices(X23,X26,X25))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_filter_0])])])])).
cnf(c_0_20, plain, (v2_struct_0(X1)|m1_subset_1(k4_lattices(X1,X2,X3),u1_struct_0(X1))|~v6_lattices(X1)|~l1_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (v6_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (l1_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (v2_struct_0(X1)|m1_subset_1(k3_lattices(X1,X2,X3),u1_struct_0(X1))|~v4_lattices(X1)|~l2_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (v4_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_25, plain, (l2_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_26, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v3_filter_0(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_lattices(X1,k4_filter_0(X1,X2,X3),k4_filter_0(X1,X2,k3_lattices(X1,X3,X4)))&(r3_lattices(X1,k4_filter_0(X1,X2,X3),k4_filter_0(X1,k4_lattices(X1,X2,X4),X3))&(r3_lattices(X1,k4_filter_0(X1,X2,X3),k4_filter_0(X1,X2,k3_lattices(X1,X4,X3)))&r3_lattices(X1,k4_filter_0(X1,X2,X3),k4_filter_0(X1,k4_lattices(X1,X4,X2),X3)))))))))), inference(assume_negation,[status(cth)],[t56_filter_1])).
fof(c_0_27, plain, ![X31, X32, X33, X34]:(v2_struct_0(X31)|(~v10_lattices(X31)|~l3_lattices(X31))|(~m1_subset_1(X32,u1_struct_0(X31))|(~m1_subset_1(X33,u1_struct_0(X31))|(~m1_subset_1(X34,u1_struct_0(X31))|(~r3_lattices(X31,X32,X33)|r3_lattices(X31,k4_lattices(X31,X32,X34),X33)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t2_filter_0])])])])).
cnf(c_0_28, plain, (v2_struct_0(X1)|k4_lattices(X1,X2,X3)=k4_lattices(X1,X3,X2)|~v6_lattices(X1)|~l1_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_29, plain, (v2_struct_0(X2)|r3_lattices(X2,X1,X5)|X5!=k4_filter_0(X2,X3,X4)|~v10_lattices(X2)|~l3_lattices(X2)|~v3_filter_0(X2)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~r3_lattices(X2,k4_lattices(X2,X3,X1),X4)), inference(cn,[status(thm)],[c_0_18])).
cnf(c_0_30, plain, (v2_struct_0(X1)|r3_lattices(X1,X2,k3_lattices(X1,X4,X3))|~v10_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r3_lattices(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_31, plain, (m1_subset_1(k4_lattices(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])).
cnf(c_0_32, plain, (m1_subset_1(k3_lattices(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
cnf(c_0_33, plain, (r3_lattices(X1,k4_lattices(X1,X2,X3),X4)|v2_struct_0(X1)|v2_struct_0(X1)|X3!=k4_filter_0(X1,X2,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~v10_lattices(X1)|~v3_filter_0(X1)|~l3_lattices(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_34, negated_conjecture, ((~v2_struct_0(esk2_0)&(v10_lattices(esk2_0)&(v3_filter_0(esk2_0)&l3_lattices(esk2_0))))&(m1_subset_1(esk3_0,u1_struct_0(esk2_0))&(m1_subset_1(esk4_0,u1_struct_0(esk2_0))&(m1_subset_1(esk5_0,u1_struct_0(esk2_0))&(~r3_lattices(esk2_0,k4_filter_0(esk2_0,esk3_0,esk4_0),k4_filter_0(esk2_0,esk3_0,k3_lattices(esk2_0,esk4_0,esk5_0)))|(~r3_lattices(esk2_0,k4_filter_0(esk2_0,esk3_0,esk4_0),k4_filter_0(esk2_0,k4_lattices(esk2_0,esk3_0,esk5_0),esk4_0))|(~r3_lattices(esk2_0,k4_filter_0(esk2_0,esk3_0,esk4_0),k4_filter_0(esk2_0,esk3_0,k3_lattices(esk2_0,esk5_0,esk4_0)))|~r3_lattices(esk2_0,k4_filter_0(esk2_0,esk3_0,esk4_0),k4_filter_0(esk2_0,k4_lattices(esk2_0,esk5_0,esk3_0),esk4_0))))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_26])])])])).
cnf(c_0_35, plain, (v2_struct_0(X1)|r3_lattices(X1,k4_lattices(X1,X2,X4),X3)|~v10_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r3_lattices(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36, plain, (k4_lattices(X1,X2,X3)=k4_lattices(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_21]), c_0_22])).
fof(c_0_37, plain, ![X35, X36, X37]:(v2_struct_0(X35)|(~v4_lattices(X35)|~l2_lattices(X35))|(~m1_subset_1(X36,u1_struct_0(X35))|~m1_subset_1(X37,u1_struct_0(X35)))|k3_lattices(X35,X36,X37)=k3_lattices(X35,X37,X36)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[commutativity_k3_lattices])])])).
cnf(c_0_38, plain, (r3_lattices(X1,X2,X3)|v2_struct_0(X1)|X3!=k4_filter_0(X1,X4,k3_lattices(X1,X5,X6))|~r3_lattices(X1,k4_lattices(X1,X4,X2),X6)|~v3_filter_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_32])).
cnf(c_0_39, plain, (v2_struct_0(X1)|r3_lattices(X1,k4_lattices(X1,X2,X3),X4)|X3!=k4_filter_0(X1,X2,X4)|~v10_lattices(X1)|~l3_lattices(X1)|~v3_filter_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_33])).
cnf(c_0_40, negated_conjecture, (v3_filter_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_41, negated_conjecture, (l3_lattices(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42, negated_conjecture, (v10_lattices(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_44, plain, (r3_lattices(X1,k4_lattices(X1,X2,X3),X4)|v2_struct_0(X1)|~r3_lattices(X1,X3,X4)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
fof(c_0_45, plain, ![X38, X39, X40]:(v2_struct_0(X38)|(~v10_lattices(X38)|~l3_lattices(X38))|(~m1_subset_1(X39,u1_struct_0(X38))|~m1_subset_1(X40,u1_struct_0(X38)))|m1_subset_1(k4_filter_0(X38,X39,X40),u1_struct_0(X38))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_filter_0])])])).
fof(c_0_46, plain, ![X41, X42, X43, X44]:(v2_struct_0(X41)|(~v10_lattices(X41)|(~v3_filter_0(X41)|~l3_lattices(X41)))|(~m1_subset_1(X42,u1_struct_0(X41))|(~m1_subset_1(X43,u1_struct_0(X41))|(~m1_subset_1(X44,u1_struct_0(X41))|k4_filter_0(X41,k4_lattices(X41,X42,X43),X44)=k4_filter_0(X41,X42,k4_filter_0(X41,X43,X44)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_filter_1])])])])).
cnf(c_0_47, plain, (v2_struct_0(X1)|k3_lattices(X1,X2,X3)=k3_lattices(X1,X3,X2)|~v4_lattices(X1)|~l2_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_48, plain, (r3_lattices(X1,X2,k4_filter_0(X1,X3,k3_lattices(X1,X4,X5)))|v2_struct_0(X1)|~r3_lattices(X1,k4_lattices(X1,X3,X2),X5)|~v3_filter_0(X1)|~m1_subset_1(k4_filter_0(X1,X3,k3_lattices(X1,X4,X5)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_49, negated_conjecture, (r3_lattices(esk2_0,k4_lattices(esk2_0,X1,X2),X3)|X2!=k4_filter_0(esk2_0,X1,X3)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41]), c_0_42])]), c_0_43])).
cnf(c_0_50, plain, (r3_lattices(X1,X2,X3)|v2_struct_0(X1)|X3!=k4_filter_0(X1,X4,X5)|~r3_lattices(X1,X2,X5)|~v3_filter_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_29, c_0_44])).
cnf(c_0_51, plain, (v2_struct_0(X1)|m1_subset_1(k4_filter_0(X1,X2,X3),u1_struct_0(X1))|~v10_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_52, plain, (v2_struct_0(X1)|k4_filter_0(X1,k4_lattices(X1,X2,X3),X4)=k4_filter_0(X1,X2,k4_filter_0(X1,X3,X4))|~v10_lattices(X1)|~v3_filter_0(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_53, negated_conjecture, (~r3_lattices(esk2_0,k4_filter_0(esk2_0,esk3_0,esk4_0),k4_filter_0(esk2_0,esk3_0,k3_lattices(esk2_0,esk4_0,esk5_0)))|~r3_lattices(esk2_0,k4_filter_0(esk2_0,esk3_0,esk4_0),k4_filter_0(esk2_0,k4_lattices(esk2_0,esk3_0,esk5_0),esk4_0))|~r3_lattices(esk2_0,k4_filter_0(esk2_0,esk3_0,esk4_0),k4_filter_0(esk2_0,esk3_0,k3_lattices(esk2_0,esk5_0,esk4_0)))|~r3_lattices(esk2_0,k4_filter_0(esk2_0,esk3_0,esk4_0),k4_filter_0(esk2_0,k4_lattices(esk2_0,esk5_0,esk3_0),esk4_0))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_54, plain, (k3_lattices(X1,X2,X3)=k3_lattices(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_24]), c_0_25])).
cnf(c_0_55, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_56, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_57, negated_conjecture, (r3_lattices(esk2_0,X1,k4_filter_0(esk2_0,X2,k3_lattices(esk2_0,X3,X4)))|X1!=k4_filter_0(esk2_0,X2,X4)|~m1_subset_1(k4_filter_0(esk2_0,X2,k3_lattices(esk2_0,X3,X4)),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_40]), c_0_41]), c_0_42])]), c_0_43])).
cnf(c_0_58, plain, (r3_lattices(X1,X2,k4_filter_0(X1,X3,X4))|v2_struct_0(X1)|~r3_lattices(X1,X2,X4)|~v3_filter_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_50]), c_0_51])).
cnf(c_0_59, negated_conjecture, (m1_subset_1(k4_filter_0(esk2_0,X1,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_42]), c_0_41])]), c_0_43])).
cnf(c_0_60, negated_conjecture, (k4_filter_0(esk2_0,k4_lattices(esk2_0,X1,X2),X3)=k4_filter_0(esk2_0,X1,k4_filter_0(esk2_0,X2,X3))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_40]), c_0_41]), c_0_42])]), c_0_43])).
cnf(c_0_61, negated_conjecture, (~r3_lattices(esk2_0,k4_filter_0(esk2_0,esk3_0,esk4_0),k4_filter_0(esk2_0,esk3_0,k3_lattices(esk2_0,esk4_0,esk5_0)))|~r3_lattices(esk2_0,k4_filter_0(esk2_0,esk3_0,esk4_0),k4_filter_0(esk2_0,k4_lattices(esk2_0,esk3_0,esk5_0),esk4_0))|~r3_lattices(esk2_0,k4_filter_0(esk2_0,esk3_0,esk4_0),k4_filter_0(esk2_0,k4_lattices(esk2_0,esk5_0,esk3_0),esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55]), c_0_56]), c_0_41]), c_0_42])]), c_0_43])).
cnf(c_0_62, negated_conjecture, (r3_lattices(esk2_0,X1,k4_filter_0(esk2_0,X2,k3_lattices(esk2_0,X3,X4)))|X1!=k4_filter_0(esk2_0,X2,X3)|~m1_subset_1(k4_filter_0(esk2_0,X2,k3_lattices(esk2_0,X3,X4)),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_54]), c_0_41]), c_0_42])]), c_0_43])).
cnf(c_0_63, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_64, negated_conjecture, (r3_lattices(esk2_0,X1,k4_filter_0(esk2_0,X2,X3))|~r3_lattices(esk2_0,X1,X3)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_40]), c_0_41]), c_0_42])]), c_0_43])).
cnf(c_0_65, negated_conjecture, (m1_subset_1(k4_filter_0(esk2_0,X1,k4_filter_0(esk2_0,X2,X3)),u1_struct_0(esk2_0))|~m1_subset_1(k4_lattices(esk2_0,X1,X2),u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_66, negated_conjecture, (~r3_lattices(esk2_0,k4_filter_0(esk2_0,esk3_0,esk4_0),k4_filter_0(esk2_0,k4_lattices(esk2_0,esk3_0,esk5_0),esk4_0))|~r3_lattices(esk2_0,k4_filter_0(esk2_0,esk3_0,esk4_0),k4_filter_0(esk2_0,k4_lattices(esk2_0,esk5_0,esk3_0),esk4_0))|~m1_subset_1(k4_filter_0(esk2_0,esk3_0,k3_lattices(esk2_0,esk4_0,esk5_0)),u1_struct_0(esk2_0))|~m1_subset_1(k4_filter_0(esk2_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63]), c_0_55]), c_0_56])])).
cnf(c_0_67, negated_conjecture, (r3_lattices(esk2_0,X1,X2)|X2!=k4_filter_0(esk2_0,X3,k4_filter_0(esk2_0,X4,X5))|~r3_lattices(esk2_0,k4_lattices(esk2_0,X3,X1),X5)|~m1_subset_1(k4_lattices(esk2_0,X3,X1),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X5,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_64]), c_0_40]), c_0_41]), c_0_42])]), c_0_43]), c_0_59])).
cnf(c_0_68, negated_conjecture, (m1_subset_1(k4_filter_0(esk2_0,X1,k4_filter_0(esk2_0,X2,X3)),u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_31]), c_0_41]), c_0_42])]), c_0_43])).
cnf(c_0_69, negated_conjecture, (~r3_lattices(esk2_0,k4_filter_0(esk2_0,esk3_0,esk4_0),k4_filter_0(esk2_0,k4_lattices(esk2_0,esk3_0,esk5_0),esk4_0))|~r3_lattices(esk2_0,k4_filter_0(esk2_0,esk3_0,esk4_0),k4_filter_0(esk2_0,k4_lattices(esk2_0,esk5_0,esk3_0),esk4_0))|~m1_subset_1(k4_filter_0(esk2_0,esk3_0,esk4_0),u1_struct_0(esk2_0))|~m1_subset_1(k3_lattices(esk2_0,esk4_0,esk5_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_59]), c_0_63])])).
cnf(c_0_70, negated_conjecture, (r3_lattices(esk2_0,X1,k4_filter_0(esk2_0,X2,k4_filter_0(esk2_0,X3,X4)))|~r3_lattices(esk2_0,k4_lattices(esk2_0,X2,X1),X4)|~m1_subset_1(k4_lattices(esk2_0,X2,X1),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_67]), c_0_68])).
cnf(c_0_71, negated_conjecture, (~r3_lattices(esk2_0,k4_filter_0(esk2_0,esk3_0,esk4_0),k4_filter_0(esk2_0,k4_lattices(esk2_0,esk3_0,esk5_0),esk4_0))|~r3_lattices(esk2_0,k4_filter_0(esk2_0,esk3_0,esk4_0),k4_filter_0(esk2_0,k4_lattices(esk2_0,esk5_0,esk3_0),esk4_0))|~m1_subset_1(k4_filter_0(esk2_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_32]), c_0_55]), c_0_56]), c_0_41]), c_0_42])]), c_0_43])).
cnf(c_0_72, negated_conjecture, (r3_lattices(esk2_0,X1,k4_filter_0(esk2_0,X2,k4_filter_0(esk2_0,X3,X4)))|X1!=k4_filter_0(esk2_0,X2,X4)|~m1_subset_1(k4_lattices(esk2_0,X2,X1),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_70, c_0_49])).
cnf(c_0_73, negated_conjecture, (k4_filter_0(esk2_0,k4_lattices(esk2_0,X1,X2),X3)=k4_filter_0(esk2_0,X2,k4_filter_0(esk2_0,X1,X3))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_36]), c_0_41]), c_0_42])]), c_0_43])).
cnf(c_0_74, negated_conjecture, (~r3_lattices(esk2_0,k4_filter_0(esk2_0,esk3_0,esk4_0),k4_filter_0(esk2_0,k4_lattices(esk2_0,esk3_0,esk5_0),esk4_0))|~m1_subset_1(k4_filter_0(esk2_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_36]), c_0_63]), c_0_55]), c_0_41]), c_0_42])]), c_0_43])).
cnf(c_0_75, negated_conjecture, (r3_lattices(esk2_0,X1,k4_filter_0(esk2_0,X2,k4_filter_0(esk2_0,X3,X4)))|X1!=k4_filter_0(esk2_0,X2,X4)|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_31]), c_0_41]), c_0_42])]), c_0_43])).
cnf(c_0_76, negated_conjecture, (k4_filter_0(esk2_0,X1,k4_filter_0(esk2_0,X2,X3))=k4_filter_0(esk2_0,X2,k4_filter_0(esk2_0,X1,X3))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_60, c_0_73])).
cnf(c_0_77, negated_conjecture, (~r3_lattices(esk2_0,k4_filter_0(esk2_0,esk3_0,esk4_0),k4_filter_0(esk2_0,esk5_0,k4_filter_0(esk2_0,esk3_0,esk4_0)))|~m1_subset_1(k4_filter_0(esk2_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_73]), c_0_56]), c_0_63]), c_0_55])])).
cnf(c_0_78, negated_conjecture, (r3_lattices(esk2_0,X1,k4_filter_0(esk2_0,X2,k4_filter_0(esk2_0,X3,X4)))|X1!=k4_filter_0(esk2_0,X3,X4)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_79, negated_conjecture, (~m1_subset_1(k4_filter_0(esk2_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_63]), c_0_56]), c_0_55])])).
cnf(c_0_80, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_59]), c_0_56]), c_0_63])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 81
# Proof object clause steps            : 56
# Proof object formula steps           : 25
# Proof object conjectures             : 34
# Proof object clause conjectures      : 31
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 32
# Proof object simplifying inferences  : 87
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 29
# Processed clauses                    : 1892
# ...of these trivial                  : 0
# ...subsumed                          : 1432
# ...remaining for further processing  : 460
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 111
# Backward-rewritten                   : 0
# Generated clauses                    : 17356
# ...of the previous two non-trivial   : 17354
# Contextual simplify-reflections      : 96
# Paramodulations                      : 17240
# Factorizations                       : 0
# Equation resolutions                 : 116
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 320
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 312
# Current number of unprocessed clauses: 15402
# ...number of literals in the above   : 164626
# Current number of archived formulas  : 0
# Current number of archived clauses   : 140
# Clause-clause subsumption calls (NU) : 82392
# Rec. Clause-clause subsumption calls : 42263
# Non-unit clause-clause subsumptions  : 1639
# Unit Clause-clause subsumption calls : 228
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 862995

# -------------------------------------------------
# User time                : 0.338 s
# System time              : 0.003 s
# Total time               : 0.341 s
# Maximum resident set size: 28700 pages
