# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_enumset1, axiom, ![X1, X2, X3, X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3)))))), file('data/MPTP/proofs/7391e1799691439e837a6189fba499b6.E_input', d1_enumset1)).
fof(t25_setfam_1, conjecture, ![X1, X2, X3]:(r1_setfam_1(X3,k2_tarski(X1,X2))=>![X4]:~((r2_hidden(X4,X3)&(~(r1_tarski(X4,X1))&~(r1_tarski(X4,X2)))))), file('data/MPTP/proofs/7391e1799691439e837a6189fba499b6.E_input', t25_setfam_1)).
fof(d2_setfam_1, axiom, ![X1, X2]:(r1_setfam_1(X1,X2)<=>![X3]:~((r2_hidden(X3,X1)&![X4]:~((r2_hidden(X4,X2)&r1_tarski(X3,X4)))))), file('data/MPTP/proofs/7391e1799691439e837a6189fba499b6.E_input', d2_setfam_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/7391e1799691439e837a6189fba499b6.E_input', t70_enumset1)).
fof(c_0_4, plain, ![X18, X19, X20, X21, X22, X23, X24, X25, X26, X27]:(((~r2_hidden(X22,X21)|(X22=X18|(X22=X19|X22=X20))|X21!=k1_enumset1(X18,X19,X20))&((X23!=X18|r2_hidden(X23,X21)|X21!=k1_enumset1(X18,X19,X20))&((X23!=X19|r2_hidden(X23,X21)|X21!=k1_enumset1(X18,X19,X20))&(X23!=X20|r2_hidden(X23,X21)|X21!=k1_enumset1(X18,X19,X20)))))&(((esk7_4(X24,X25,X26,X27)!=X24|~r2_hidden(esk7_4(X24,X25,X26,X27),X27)|X27=k1_enumset1(X24,X25,X26))&((esk7_4(X24,X25,X26,X27)!=X25|~r2_hidden(esk7_4(X24,X25,X26,X27),X27)|X27=k1_enumset1(X24,X25,X26))&(esk7_4(X24,X25,X26,X27)!=X26|~r2_hidden(esk7_4(X24,X25,X26,X27),X27)|X27=k1_enumset1(X24,X25,X26))))&(r2_hidden(esk7_4(X24,X25,X26,X27),X27)|(esk7_4(X24,X25,X26,X27)=X24|(esk7_4(X24,X25,X26,X27)=X25|esk7_4(X24,X25,X26,X27)=X26))|X27=k1_enumset1(X24,X25,X26)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])).
fof(c_0_5, negated_conjecture, ~(![X1, X2, X3]:(r1_setfam_1(X3,k2_tarski(X1,X2))=>![X4]:~((r2_hidden(X4,X3)&(~(r1_tarski(X4,X1))&~(r1_tarski(X4,X2))))))), inference(assume_negation,[status(cth)],[t25_setfam_1])).
cnf(c_0_6, plain, (X1=X3|X1=X4|X1=X5|~r2_hidden(X1,X2)|X2!=k1_enumset1(X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_7, plain, ![X10, X11, X12, X14, X15, X17]:(((r2_hidden(esk5_3(X10,X11,X12),X11)|~r2_hidden(X12,X10)|~r1_setfam_1(X10,X11))&(r1_tarski(X12,esk5_3(X10,X11,X12))|~r2_hidden(X12,X10)|~r1_setfam_1(X10,X11)))&((r2_hidden(esk6_2(X14,X15),X14)|r1_setfam_1(X14,X15))&(~r2_hidden(X17,X15)|~r1_tarski(esk6_2(X14,X15),X17)|r1_setfam_1(X14,X15)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_setfam_1])])])])])])).
fof(c_0_8, negated_conjecture, (r1_setfam_1(esk3_0,k2_tarski(esk1_0,esk2_0))&(r2_hidden(esk4_0,esk3_0)&(~r1_tarski(esk4_0,esk1_0)&~r1_tarski(esk4_0,esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_5])])])])).
fof(c_0_9, plain, ![X29, X30]:k1_enumset1(X29,X29,X30)=k2_tarski(X29,X30), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_10, plain, (X1=X2|X1=X3|X1=X4|~r2_hidden(X1,k1_enumset1(X4,X3,X2))), inference(er,[status(thm)],[c_0_6])).
cnf(c_0_11, plain, (r2_hidden(esk5_3(X1,X2,X3),X2)|~r2_hidden(X3,X1)|~r1_setfam_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, negated_conjecture, (r1_setfam_1(esk3_0,k2_tarski(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (esk5_3(X1,k1_enumset1(X2,X3,X4),X5)=X4|esk5_3(X1,k1_enumset1(X2,X3,X4),X5)=X3|esk5_3(X1,k1_enumset1(X2,X3,X4),X5)=X2|~r2_hidden(X5,X1)|~r1_setfam_1(X1,k1_enumset1(X2,X3,X4))), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15, negated_conjecture, (r1_setfam_1(esk3_0,k1_enumset1(esk1_0,esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_16, plain, (r1_tarski(X1,esk5_3(X2,X3,X1))|~r2_hidden(X1,X2)|~r1_setfam_1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_17, negated_conjecture, (esk5_3(esk3_0,k1_enumset1(esk1_0,esk1_0,esk2_0),X1)=esk1_0|esk5_3(esk3_0,k1_enumset1(esk1_0,esk1_0,esk2_0),X1)=esk2_0|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_18, negated_conjecture, (esk5_3(esk3_0,k1_enumset1(esk1_0,esk1_0,esk2_0),X1)=esk1_0|r1_tarski(X1,esk2_0)|~r2_hidden(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_15])])).
cnf(c_0_19, negated_conjecture, (~r1_tarski(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_20, negated_conjecture, (r1_tarski(X1,esk2_0)|r1_tarski(X1,esk1_0)|~r2_hidden(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_18]), c_0_15])])).
cnf(c_0_21, negated_conjecture, (r2_hidden(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_22, negated_conjecture, (~r1_tarski(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_23, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])]), c_0_22]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 24
# Proof object clause steps            : 15
# Proof object formula steps           : 9
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 9
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 16
# Processed clauses                    : 48
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 48
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 31
# ...of the previous two non-trivial   : 23
# Contextual simplify-reflections      : 0
# Paramodulations                      : 26
# Factorizations                       : 1
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 28
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 20
# Current number of unprocessed clauses: 6
# ...number of literals in the above   : 27
# Current number of archived formulas  : 0
# Current number of archived clauses   : 17
# Clause-clause subsumption calls (NU) : 29
# Rec. Clause-clause subsumption calls : 27
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 20
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1577

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28652 pages
