# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/7427b1b296a440dab1bbcdeeddf5317e.E_input', d8_lattice3)).
fof(t31_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3))=>(r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2)))))&((r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2))))=>(X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3)))))), file('data/MPTP/proofs/7427b1b296a440dab1bbcdeeddf5317e.E_input', t31_yellow_0)).
fof(t18_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k6_waybel_0(X1,X2))<=>r1_orders_2(X1,X2,X3))))), file('data/MPTP/proofs/7427b1b296a440dab1bbcdeeddf5317e.E_input', t18_waybel_0)).
fof(d10_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_yellow_0(X1,X2)=>(X3=k2_yellow_0(X1,X2)<=>(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))))))), file('data/MPTP/proofs/7427b1b296a440dab1bbcdeeddf5317e.E_input', d10_yellow_0)).
fof(t39_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2))), file('data/MPTP/proofs/7427b1b296a440dab1bbcdeeddf5317e.E_input', t39_yellow_0)).
fof(t38_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))&r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))))), file('data/MPTP/proofs/7427b1b296a440dab1bbcdeeddf5317e.E_input', t38_yellow_0)).
fof(t39_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_yellow_0(X1,k6_waybel_0(X1,X2))&k2_yellow_0(X1,k6_waybel_0(X1,X2))=X2))), file('data/MPTP/proofs/7427b1b296a440dab1bbcdeeddf5317e.E_input', t39_waybel_0)).
fof(c_0_7, plain, ![X5, X6, X7, X8]:((~r1_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X6)|r1_orders_2(X5,X7,X8)))|~m1_subset_1(X7,u1_struct_0(X5))|~l1_orders_2(X5))&((m1_subset_1(esk1_3(X5,X6,X7),u1_struct_0(X5))|r1_lattice3(X5,X6,X7)|~m1_subset_1(X7,u1_struct_0(X5))|~l1_orders_2(X5))&((r2_hidden(esk1_3(X5,X6,X7),X6)|r1_lattice3(X5,X6,X7)|~m1_subset_1(X7,u1_struct_0(X5))|~l1_orders_2(X5))&(~r1_orders_2(X5,X7,esk1_3(X5,X6,X7))|r1_lattice3(X5,X6,X7)|~m1_subset_1(X7,u1_struct_0(X5))|~l1_orders_2(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
fof(c_0_8, plain, ![X10, X11, X12, X13, X14]:(((r1_lattice3(X10,X12,X11)|(X11!=k2_yellow_0(X10,X12)|~r2_yellow_0(X10,X12))|~m1_subset_1(X11,u1_struct_0(X10))|(~v5_orders_2(X10)|~l1_orders_2(X10)))&(~m1_subset_1(X13,u1_struct_0(X10))|(~r1_lattice3(X10,X12,X13)|r1_orders_2(X10,X13,X11))|(X11!=k2_yellow_0(X10,X12)|~r2_yellow_0(X10,X12))|~m1_subset_1(X11,u1_struct_0(X10))|(~v5_orders_2(X10)|~l1_orders_2(X10))))&(((X11=k2_yellow_0(X10,X14)|(m1_subset_1(esk2_3(X10,X11,X14),u1_struct_0(X10))|~r1_lattice3(X10,X14,X11))|~m1_subset_1(X11,u1_struct_0(X10))|(~v5_orders_2(X10)|~l1_orders_2(X10)))&(r2_yellow_0(X10,X14)|(m1_subset_1(esk2_3(X10,X11,X14),u1_struct_0(X10))|~r1_lattice3(X10,X14,X11))|~m1_subset_1(X11,u1_struct_0(X10))|(~v5_orders_2(X10)|~l1_orders_2(X10))))&(((X11=k2_yellow_0(X10,X14)|(r1_lattice3(X10,X14,esk2_3(X10,X11,X14))|~r1_lattice3(X10,X14,X11))|~m1_subset_1(X11,u1_struct_0(X10))|(~v5_orders_2(X10)|~l1_orders_2(X10)))&(r2_yellow_0(X10,X14)|(r1_lattice3(X10,X14,esk2_3(X10,X11,X14))|~r1_lattice3(X10,X14,X11))|~m1_subset_1(X11,u1_struct_0(X10))|(~v5_orders_2(X10)|~l1_orders_2(X10))))&((X11=k2_yellow_0(X10,X14)|(~r1_orders_2(X10,esk2_3(X10,X11,X14),X11)|~r1_lattice3(X10,X14,X11))|~m1_subset_1(X11,u1_struct_0(X10))|(~v5_orders_2(X10)|~l1_orders_2(X10)))&(r2_yellow_0(X10,X14)|(~r1_orders_2(X10,esk2_3(X10,X11,X14),X11)|~r1_lattice3(X10,X14,X11))|~m1_subset_1(X11,u1_struct_0(X10))|(~v5_orders_2(X10)|~l1_orders_2(X10))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_yellow_0])])])])])])).
fof(c_0_9, plain, ![X16, X17, X18]:((~r2_hidden(X18,k6_waybel_0(X16,X17))|r1_orders_2(X16,X17,X18)|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|~l1_orders_2(X16)))&(~r1_orders_2(X16,X17,X18)|r2_hidden(X18,k6_waybel_0(X16,X17))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|~l1_orders_2(X16)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t18_waybel_0])])])])])).
fof(c_0_10, plain, ![X19, X20, X21, X22]:(((r1_lattice3(X19,X20,X21)|X21!=k2_yellow_0(X19,X20)|~r2_yellow_0(X19,X20)|~m1_subset_1(X21,u1_struct_0(X19))|~l1_orders_2(X19))&(~m1_subset_1(X22,u1_struct_0(X19))|(~r1_lattice3(X19,X20,X22)|r1_orders_2(X19,X22,X21))|X21!=k2_yellow_0(X19,X20)|~r2_yellow_0(X19,X20)|~m1_subset_1(X21,u1_struct_0(X19))|~l1_orders_2(X19)))&((m1_subset_1(esk3_3(X19,X20,X21),u1_struct_0(X19))|~r1_lattice3(X19,X20,X21)|X21=k2_yellow_0(X19,X20)|~r2_yellow_0(X19,X20)|~m1_subset_1(X21,u1_struct_0(X19))|~l1_orders_2(X19))&((r1_lattice3(X19,X20,esk3_3(X19,X20,X21))|~r1_lattice3(X19,X20,X21)|X21=k2_yellow_0(X19,X20)|~r2_yellow_0(X19,X20)|~m1_subset_1(X21,u1_struct_0(X19))|~l1_orders_2(X19))&(~r1_orders_2(X19,esk3_3(X19,X20,X21),X21)|~r1_lattice3(X19,X20,X21)|X21=k2_yellow_0(X19,X20)|~r2_yellow_0(X19,X20)|~m1_subset_1(X21,u1_struct_0(X19))|~l1_orders_2(X19))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_yellow_0])])])])])).
cnf(c_0_11, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (r2_yellow_0(X1,X2)|r1_lattice3(X1,X2,esk2_3(X1,X3,X2))|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (r2_yellow_0(X1,X2)|m1_subset_1(esk2_3(X1,X3,X2),u1_struct_0(X1))|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, plain, (r1_orders_2(X2,X3,X1)|v2_struct_0(X2)|~r2_hidden(X1,k6_waybel_0(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (r2_hidden(esk1_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_16, plain, (r1_orders_2(X2,X1,X4)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|X4!=k2_yellow_0(X2,X3)|~r2_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_17, plain, ![X24, X25]:((k1_yellow_0(X24,k6_domain_1(u1_struct_0(X24),X25))=X25|~m1_subset_1(X25,u1_struct_0(X24))|(v2_struct_0(X24)|(~v3_orders_2(X24)|(~v5_orders_2(X24)|~l1_orders_2(X24)))))&(k2_yellow_0(X24,k6_domain_1(u1_struct_0(X24),X25))=X25|~m1_subset_1(X25,u1_struct_0(X24))|(v2_struct_0(X24)|(~v3_orders_2(X24)|(~v5_orders_2(X24)|~l1_orders_2(X24)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t39_yellow_0])])])])])).
fof(c_0_18, plain, ![X26, X27]:((r1_yellow_0(X26,k6_domain_1(u1_struct_0(X26),X27))|~m1_subset_1(X27,u1_struct_0(X26))|(v2_struct_0(X26)|(~v3_orders_2(X26)|(~v5_orders_2(X26)|~l1_orders_2(X26)))))&(r2_yellow_0(X26,k6_domain_1(u1_struct_0(X26),X27))|~m1_subset_1(X27,u1_struct_0(X26))|(v2_struct_0(X26)|(~v3_orders_2(X26)|(~v5_orders_2(X26)|~l1_orders_2(X26)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t38_yellow_0])])])])])).
cnf(c_0_19, plain, (r1_lattice3(X1,X2,X3)|X3!=k2_yellow_0(X1,X2)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_20, plain, (r2_yellow_0(X1,X2)|~r1_orders_2(X1,esk2_3(X1,X3,X2),X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_21, plain, (r2_yellow_0(X1,X2)|r1_orders_2(X1,esk2_3(X1,X3,X2),X4)|~v5_orders_2(X1)|~r2_hidden(X4,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])).
cnf(c_0_22, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk1_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_23, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,esk1_3(X3,k6_waybel_0(X1,X2),X4))|r1_lattice3(X3,k6_waybel_0(X1,X2),X4)|~m1_subset_1(esk1_3(X3,k6_waybel_0(X1,X2),X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_orders_2(X1)|~l1_orders_2(X3)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_24, plain, (m1_subset_1(esk1_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_25, plain, (r1_orders_2(X1,X2,k2_yellow_0(X1,X3))|~r2_yellow_0(X1,X3)|~r1_lattice3(X1,X3,X2)|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_26, plain, (k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27, plain, (r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, plain, (r1_lattice3(X1,X2,k2_yellow_0(X1,X2))|~r2_yellow_0(X1,X2)|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_19])).
fof(c_0_29, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_yellow_0(X1,k6_waybel_0(X1,X2))&k2_yellow_0(X1,k6_waybel_0(X1,X2))=X2)))), inference(assume_negation,[status(cth)],[t39_waybel_0])).
cnf(c_0_30, plain, (X1=k2_yellow_0(X2,X3)|r1_lattice3(X2,X3,esk2_3(X2,X1,X3))|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_31, plain, (X1=k2_yellow_0(X2,X3)|m1_subset_1(esk2_3(X2,X1,X3),u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_32, plain, (r2_yellow_0(X1,X2)|~v5_orders_2(X1)|~r2_hidden(X3,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_33, plain, (v2_struct_0(X1)|r1_lattice3(X1,k6_waybel_0(X1,X2),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24])).
cnf(c_0_34, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~v5_orders_2(X1)|~r1_lattice3(X1,k6_domain_1(u1_struct_0(X1),X3),X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])).
cnf(c_0_35, plain, (v2_struct_0(X1)|r1_lattice3(X1,k6_domain_1(u1_struct_0(X1),X2),X2)|~v3_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_26]), c_0_27])).
fof(c_0_36, negated_conjecture, ((~v2_struct_0(esk4_0)&(v3_orders_2(esk4_0)&(v4_orders_2(esk4_0)&(v5_orders_2(esk4_0)&l1_orders_2(esk4_0)))))&(m1_subset_1(esk5_0,u1_struct_0(esk4_0))&(~r2_yellow_0(esk4_0,k6_waybel_0(esk4_0,esk5_0))|k2_yellow_0(esk4_0,k6_waybel_0(esk4_0,esk5_0))!=esk5_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_29])])])])).
cnf(c_0_37, plain, (X1=k2_yellow_0(X2,X3)|~r1_orders_2(X2,esk2_3(X2,X1,X3),X1)|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_38, plain, (X1=k2_yellow_0(X2,X3)|r1_orders_2(X2,esk2_3(X2,X1,X3),X4)|~v5_orders_2(X2)|~r2_hidden(X4,X3)|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_30]), c_0_31])).
cnf(c_0_39, plain, (v2_struct_0(X1)|r2_yellow_0(X1,k6_waybel_0(X1,X2))|~v5_orders_2(X1)|~r2_hidden(X2,k6_waybel_0(X1,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_40, plain, (r2_hidden(X3,k6_waybel_0(X1,X2))|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_41, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_42, negated_conjecture, (v3_orders_2(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_43, negated_conjecture, (v5_orders_2(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44, negated_conjecture, (l1_orders_2(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_45, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_46, plain, (X1=k2_yellow_0(X2,X3)|~v5_orders_2(X2)|~r2_hidden(X1,X3)|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_47, plain, (v2_struct_0(X1)|r2_yellow_0(X1,k6_waybel_0(X1,X2))|~v5_orders_2(X1)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_48, negated_conjecture, (r1_orders_2(esk4_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), c_0_44])]), c_0_45])).
cnf(c_0_49, plain, (k2_yellow_0(X1,k6_waybel_0(X1,X2))=X2|v2_struct_0(X1)|~v5_orders_2(X1)|~r2_hidden(X2,k6_waybel_0(X1,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_46, c_0_33])).
cnf(c_0_50, negated_conjecture, (~r2_yellow_0(esk4_0,k6_waybel_0(esk4_0,esk5_0))|k2_yellow_0(esk4_0,k6_waybel_0(esk4_0,esk5_0))!=esk5_0), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_51, negated_conjecture, (r2_yellow_0(esk4_0,k6_waybel_0(esk4_0,X1))|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_43]), c_0_44])]), c_0_45])).
cnf(c_0_52, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_53, plain, (k2_yellow_0(X1,k6_waybel_0(X1,X2))=X2|v2_struct_0(X1)|~v5_orders_2(X1)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_49, c_0_40])).
cnf(c_0_54, negated_conjecture, (k2_yellow_0(esk4_0,k6_waybel_0(esk4_0,esk5_0))!=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52])])).
cnf(c_0_55, negated_conjecture, (k2_yellow_0(esk4_0,k6_waybel_0(esk4_0,X1))=X1|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_48]), c_0_43]), c_0_44])]), c_0_45])).
cnf(c_0_56, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_52])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 57
# Proof object clause steps            : 42
# Proof object formula steps           : 15
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 23
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 104
# ...of these trivial                  : 0
# ...subsumed                          : 14
# ...remaining for further processing  : 90
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 81
# ...of the previous two non-trivial   : 75
# Contextual simplify-reflections      : 14
# Paramodulations                      : 77
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 57
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 50
# Current number of unprocessed clauses: 22
# ...number of literals in the above   : 195
# Current number of archived formulas  : 0
# Current number of archived clauses   : 29
# Clause-clause subsumption calls (NU) : 522
# Rec. Clause-clause subsumption calls : 49
# Non-unit clause-clause subsumptions  : 28
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5802

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.000 s
# Total time               : 0.013 s
# Maximum resident set size: 28792 pages
