# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', t70_enumset1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', t7_xboole_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', commutativity_k2_tarski)).
fof(t44_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k4_xboole_0(X1,X2),X3)=>r1_tarski(X1,k2_xboole_0(X2,X3))), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', t44_xboole_1)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', redefinition_k6_subset_1)).
fof(idempotence_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', idempotence_k2_xboole_0)).
fof(t31_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>r1_tarski(k3_relat_1(X1),k3_relat_1(X2))))), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', t31_relat_1)).
fof(t43_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3)), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', t43_xboole_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', t3_xboole_0)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', t65_xboole_1)).
fof(t10_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(X1,k2_xboole_0(X3,X2))), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', t10_xboole_1)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', t1_xboole_1)).
fof(t7_tarski, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&r2_hidden(X4,X3))))))), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', t7_tarski)).
fof(t9_tarski, axiom, ![X1]:?[X2]:(r2_hidden(X1,X2)&(![X3, X4]:((r2_hidden(X3,X2)&r1_tarski(X4,X3))=>r2_hidden(X4,X2))&(![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&![X5]:(r1_tarski(X5,X3)=>r2_hidden(X5,X4))))))&![X3]:~((r1_tarski(X3,X2)&(~(r2_tarski(X3,X2))&~(r2_hidden(X3,X2)))))))), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', t9_tarski)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', d4_relat_1)).
fof(t8_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X3,X2))=>r1_tarski(k2_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', t8_xboole_1)).
fof(d6_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k3_relat_1(X1)=k2_xboole_0(k1_relat_1(X1),k2_relat_1(X1))), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', d6_relat_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', t7_xboole_0)).
fof(cc1_relat_1, axiom, ![X1]:(v1_xboole_0(X1)=>v1_relat_1(X1)), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', cc1_relat_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', d10_xboole_0)).
fof(t15_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>r1_tarski(k6_subset_1(k1_relat_1(X1),k1_relat_1(X2)),k1_relat_1(k6_subset_1(X1,X2))))), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', t15_relat_1)).
fof(t3_xboole_1, axiom, ![X1]:(r1_tarski(X1,k1_xboole_0)=>X1=k1_xboole_0), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', t3_xboole_1)).
fof(t19_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>~((r2_hidden(X1,k2_relat_1(X2))&![X3]:~(r2_hidden(X3,k1_relat_1(X2)))))), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', t19_relat_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', fc1_xboole_0)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', t2_xboole_1)).
fof(t28_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>r1_tarski(k6_subset_1(k2_relat_1(X1),k2_relat_1(X2)),k2_relat_1(k6_subset_1(X1,X2))))), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', t28_relat_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/74304573838e49378530d47bbed7d4b4.E_input', l32_xboole_1)).
fof(c_0_28, plain, ![X62, X63]:k3_tarski(k2_tarski(X62,X63))=k2_xboole_0(X62,X63), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_29, plain, ![X75, X76]:k1_enumset1(X75,X75,X76)=k2_tarski(X75,X76), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_30, plain, ![X53, X54]:r1_tarski(X53,k2_xboole_0(X53,X54)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_31, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_32, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_33, plain, ![X77, X78]:k2_tarski(X77,X78)=k2_tarski(X78,X77), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
fof(c_0_34, plain, ![X36, X37, X38]:(~r1_tarski(k4_xboole_0(X36,X37),X38)|r1_tarski(X36,k2_xboole_0(X37,X38))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t44_xboole_1])])).
fof(c_0_35, plain, ![X28, X29]:k6_subset_1(X28,X29)=k4_xboole_0(X28,X29), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
fof(c_0_36, plain, ![X35]:k2_xboole_0(X35,X35)=X35, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k2_xboole_0])])).
fof(c_0_37, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>r1_tarski(k3_relat_1(X1),k3_relat_1(X2)))))), inference(assume_negation,[status(cth)],[t31_relat_1])).
fof(c_0_38, plain, ![X68, X69, X70]:(~r1_tarski(X68,k2_xboole_0(X69,X70))|r1_tarski(k4_xboole_0(X68,X69),X70)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
cnf(c_0_39, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_41, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_42, plain, ![X55, X56, X58, X59, X60]:(((r2_hidden(esk9_2(X55,X56),X55)|r1_xboole_0(X55,X56))&(r2_hidden(esk9_2(X55,X56),X56)|r1_xboole_0(X55,X56)))&(~r2_hidden(X60,X58)|~r2_hidden(X60,X59)|~r1_xboole_0(X58,X59))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_43, plain, ![X52]:r1_xboole_0(X52,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_44, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(k4_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_46, plain, (k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_47, plain, ![X71, X72, X73]:(~r1_tarski(X71,X72)|r1_tarski(X71,k2_xboole_0(X73,X72))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_xboole_1])])).
fof(c_0_48, plain, ![X47, X48, X49]:(~r1_tarski(X47,X48)|~r1_tarski(X48,X49)|r1_tarski(X47,X49)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_49, negated_conjecture, (v1_relat_1(esk1_0)&(v1_relat_1(esk2_0)&(r1_tarski(esk1_0,esk2_0)&~r1_tarski(k3_relat_1(esk1_0),k3_relat_1(esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_37])])])).
cnf(c_0_50, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_51, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_52, plain, (k1_enumset1(X1,X1,X2)=k1_enumset1(X2,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_32]), c_0_32])).
cnf(c_0_53, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_54, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_55, plain, ![X10, X11, X13]:((r2_hidden(esk3_2(X10,X11),X11)|~r2_hidden(X10,X11))&(~r2_hidden(X13,X11)|~r2_hidden(X13,esk3_2(X10,X11))|~r2_hidden(X10,X11))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tarski])])])])])).
cnf(c_0_56, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X2,X2,X3)))|~r1_tarski(k6_subset_1(X1,X2),X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_45]), c_0_40])).
cnf(c_0_57, plain, (k3_tarski(k1_enumset1(X1,X1,X1))=X1), inference(rw,[status(thm)],[c_0_46, c_0_40])).
cnf(c_0_58, plain, (r1_tarski(X1,k2_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_59, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_60, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_61, plain, (r1_tarski(k6_subset_1(X1,X2),X3)|~r1_tarski(X1,k3_tarski(k1_enumset1(X2,X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_45]), c_0_40])).
cnf(c_0_62, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X2,X2,X1)))), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_63, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_64, plain, (r2_hidden(esk3_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_55])).
fof(c_0_65, plain, ![X39, X41, X42, X43, X45, X46]:(r2_hidden(X39,esk7_1(X39))&((~r2_hidden(X41,esk7_1(X39))|~r1_tarski(X42,X41)|r2_hidden(X42,esk7_1(X39)))&(((r2_hidden(esk8_2(X39,X43),esk7_1(X39))|~r2_hidden(X43,esk7_1(X39)))&(~r1_tarski(X45,X43)|r2_hidden(X45,esk8_2(X39,X43))|~r2_hidden(X43,esk7_1(X39))))&(~r1_tarski(X46,esk7_1(X39))|(r2_tarski(X46,esk7_1(X39))|r2_hidden(X46,esk7_1(X39))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_tarski])])])])])])).
fof(c_0_66, plain, ![X16, X17, X18, X20, X21, X22, X23, X25]:(((~r2_hidden(X18,X17)|r2_hidden(k4_tarski(X18,esk4_3(X16,X17,X18)),X16)|X17!=k1_relat_1(X16))&(~r2_hidden(k4_tarski(X20,X21),X16)|r2_hidden(X20,X17)|X17!=k1_relat_1(X16)))&((~r2_hidden(esk5_2(X22,X23),X23)|~r2_hidden(k4_tarski(esk5_2(X22,X23),X25),X22)|X23=k1_relat_1(X22))&(r2_hidden(esk5_2(X22,X23),X23)|r2_hidden(k4_tarski(esk5_2(X22,X23),esk6_2(X22,X23)),X22)|X23=k1_relat_1(X22)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
cnf(c_0_67, plain, (r1_tarski(X1,X2)|~r1_tarski(k6_subset_1(X1,X2),X2)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_68, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X3,X3,X2)))|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_58, c_0_40])).
cnf(c_0_69, negated_conjecture, (r1_tarski(X1,esk2_0)|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_70, plain, (r1_tarski(k6_subset_1(X1,X2),X1)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_71, plain, (~r2_hidden(esk3_2(X1,k1_xboole_0),X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_72, plain, (r2_hidden(X1,esk7_1(X1))), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_73, plain, (r2_hidden(k4_tarski(X1,esk4_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_66])).
fof(c_0_74, plain, ![X32, X33, X34]:(~r1_tarski(X32,X33)|~r1_tarski(X34,X33)|r1_tarski(k2_xboole_0(X32,X34),X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_xboole_1])])).
fof(c_0_75, plain, ![X74]:(~v1_relat_1(X74)|k3_relat_1(X74)=k2_xboole_0(k1_relat_1(X74),k2_relat_1(X74))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_relat_1])])).
cnf(c_0_76, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X2,X2,X3)))|~r1_tarski(k6_subset_1(X1,k3_tarski(k1_enumset1(X2,X2,X3))),X3)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_77, negated_conjecture, (r1_tarski(k6_subset_1(esk1_0,X1),esk2_0)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_78, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_79, plain, (r2_hidden(k4_tarski(X1,esk4_3(X2,k1_relat_1(X2),X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_73])).
fof(c_0_80, plain, ![X79]:(X79=k1_xboole_0|r2_hidden(esk11_1(X79),X79)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
fof(c_0_81, plain, ![X61]:(~v1_xboole_0(X61)|v1_relat_1(X61)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relat_1])])).
cnf(c_0_82, plain, (r1_tarski(k2_xboole_0(X1,X3),X2)|~r1_tarski(X1,X2)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_83, plain, (k3_relat_1(X1)=k2_xboole_0(k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
fof(c_0_84, plain, ![X14, X15]:(((r1_tarski(X14,X15)|X14!=X15)&(r1_tarski(X15,X14)|X14!=X15))&(~r1_tarski(X14,X15)|~r1_tarski(X15,X14)|X14=X15)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_85, plain, ![X8, X9]:(~v1_relat_1(X8)|(~v1_relat_1(X9)|r1_tarski(k6_subset_1(k1_relat_1(X8),k1_relat_1(X9)),k1_relat_1(k6_subset_1(X8,X9))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_relat_1])])])).
fof(c_0_86, plain, ![X27]:(~r1_tarski(X27,k1_xboole_0)|X27=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_xboole_1])])).
cnf(c_0_87, plain, (r1_tarski(k6_subset_1(X1,X2),X3)|~r1_tarski(X1,k3_tarski(k1_enumset1(X3,X3,X2)))), inference(spm,[status(thm)],[c_0_61, c_0_52])).
cnf(c_0_88, negated_conjecture, (r1_tarski(esk1_0,k3_tarski(k1_enumset1(X1,X1,esk2_0)))), inference(spm,[status(thm)],[c_0_76, c_0_77])).
fof(c_0_89, plain, ![X64, X65]:(~v1_relat_1(X65)|(~r2_hidden(X64,k2_relat_1(X65))|r2_hidden(esk10_2(X64,X65),k1_relat_1(X65)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t19_relat_1])])])])).
cnf(c_0_90, plain, (~r2_hidden(X1,k1_relat_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_91, plain, (X1=k1_xboole_0|r2_hidden(esk11_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_92, plain, (v1_relat_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_93, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_94, plain, (r1_tarski(k3_tarski(k1_enumset1(X1,X1,X3)),X2)|~r1_tarski(X3,X2)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_82, c_0_40])).
cnf(c_0_95, plain, (k3_relat_1(X1)=k3_tarski(k1_enumset1(k1_relat_1(X1),k1_relat_1(X1),k2_relat_1(X1)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_83, c_0_40])).
cnf(c_0_96, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_97, plain, (r1_tarski(k6_subset_1(k1_relat_1(X1),k1_relat_1(X2)),k1_relat_1(k6_subset_1(X1,X2)))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_98, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_99, negated_conjecture, (r1_tarski(k6_subset_1(esk1_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
fof(c_0_100, plain, ![X67]:r1_tarski(k1_xboole_0,X67), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_101, plain, ![X50, X51]:(~v1_relat_1(X50)|(~v1_relat_1(X51)|r1_tarski(k6_subset_1(k2_relat_1(X50),k2_relat_1(X51)),k2_relat_1(k6_subset_1(X50,X51))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_relat_1])])])).
cnf(c_0_102, plain, (r2_hidden(esk10_2(X2,X1),k1_relat_1(X1))|~v1_relat_1(X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_103, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_104, plain, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_105, negated_conjecture, (~r1_tarski(k3_relat_1(esk1_0),k3_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_106, plain, (r1_tarski(k3_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X2)|~r1_tarski(k1_relat_1(X1),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_107, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_108, plain, (k6_subset_1(k1_relat_1(X1),k1_relat_1(X2))=k1_relat_1(k6_subset_1(X1,X2))|~r1_tarski(k1_relat_1(k6_subset_1(X1,X2)),k6_subset_1(k1_relat_1(X1),k1_relat_1(X2)))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_109, negated_conjecture, (k6_subset_1(esk1_0,esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_110, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_111, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_112, plain, ![X30, X31]:((k4_xboole_0(X30,X31)!=k1_xboole_0|r1_tarski(X30,X31))&(~r1_tarski(X30,X31)|k4_xboole_0(X30,X31)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_113, plain, (r1_tarski(k6_subset_1(k2_relat_1(X1),k2_relat_1(X2)),k2_relat_1(k6_subset_1(X1,X2)))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_114, plain, (~r2_hidden(X1,k2_relat_1(k1_xboole_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_104])]), c_0_78])).
cnf(c_0_115, negated_conjecture, (~r1_tarski(k2_relat_1(esk1_0),k3_relat_1(esk2_0))|~r1_tarski(k1_relat_1(esk1_0),k3_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_107])])).
cnf(c_0_116, plain, (r1_tarski(X1,k3_relat_1(X2))|~r1_tarski(X1,k2_relat_1(X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_68, c_0_95])).
cnf(c_0_117, plain, (r1_tarski(X1,k3_relat_1(X2))|~r1_tarski(k6_subset_1(X1,k1_relat_1(X2)),k2_relat_1(X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_56, c_0_95])).
cnf(c_0_118, negated_conjecture, (k6_subset_1(k1_relat_1(esk1_0),k1_relat_1(esk2_0))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_103]), c_0_103]), c_0_110]), c_0_111]), c_0_107])])).
cnf(c_0_119, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_120, plain, (k6_subset_1(k2_relat_1(X1),k2_relat_1(X2))=k2_relat_1(k6_subset_1(X1,X2))|~r1_tarski(k2_relat_1(k6_subset_1(X1,X2)),k6_subset_1(k2_relat_1(X1),k2_relat_1(X2)))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_96, c_0_113])).
cnf(c_0_121, plain, (k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_114, c_0_91])).
cnf(c_0_122, negated_conjecture, (~r1_tarski(k1_relat_1(esk1_0),k3_relat_1(esk2_0))|~r1_tarski(k2_relat_1(esk1_0),k2_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_111])])).
cnf(c_0_123, negated_conjecture, (r1_tarski(k1_relat_1(esk1_0),k3_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_110]), c_0_111])])).
cnf(c_0_124, plain, (r1_tarski(X1,X2)|k6_subset_1(X1,X2)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_119, c_0_45])).
cnf(c_0_125, negated_conjecture, (k6_subset_1(k2_relat_1(esk1_0),k2_relat_1(esk2_0))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_109]), c_0_121]), c_0_121]), c_0_110]), c_0_111]), c_0_107])])).
cnf(c_0_126, negated_conjecture, (~r1_tarski(k2_relat_1(esk1_0),k2_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_122, c_0_123])])).
cnf(c_0_127, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125]), c_0_126]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 128
# Proof object clause steps            : 72
# Proof object formula steps           : 56
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 28
# Proof object generating inferences   : 29
# Proof object simplifying inferences  : 39
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 28
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 44
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 41
# Processed clauses                    : 1134
# ...of these trivial                  : 4
# ...subsumed                          : 739
# ...remaining for further processing  : 391
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 30
# Backward-rewritten                   : 16
# Generated clauses                    : 3753
# ...of the previous two non-trivial   : 3270
# Contextual simplify-reflections      : 6
# Paramodulations                      : 3749
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 301
#    Positive orientable unit clauses  : 32
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 259
# Current number of unprocessed clauses: 2080
# ...number of literals in the above   : 7106
# Current number of archived formulas  : 0
# Current number of archived clauses   : 89
# Clause-clause subsumption calls (NU) : 8608
# Rec. Clause-clause subsumption calls : 6250
# Non-unit clause-clause subsumptions  : 515
# Unit Clause-clause subsumption calls : 533
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 33
# BW rewrite match successes           : 18
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 52348

# -------------------------------------------------
# User time                : 0.045 s
# System time              : 0.000 s
# Total time               : 0.045 s
# Maximum resident set size: 28856 pages
