# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>(r2_hidden(X1,k1_relat_1(X3))&X2=k1_funct_1(X3,X1)))), file('data/MPTP/proofs/747d1891215e47279b20aad99c2a9a2d.E_input', t8_funct_1)).
fof(d10_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(X2=k6_relat_1(X1)<=>![X3, X4]:(r2_hidden(k4_tarski(X3,X4),X2)<=>(r2_hidden(X3,X1)&X3=X4)))), file('data/MPTP/proofs/747d1891215e47279b20aad99c2a9a2d.E_input', d10_relat_1)).
fof(t34_funct_1, conjecture, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(X2=k6_relat_1(X1)<=>(k1_relat_1(X2)=X1&![X3]:(r2_hidden(X3,X1)=>k1_funct_1(X2,X3)=X3)))), file('data/MPTP/proofs/747d1891215e47279b20aad99c2a9a2d.E_input', t34_funct_1)).
fof(fc3_funct_1, axiom, ![X1]:(v1_relat_1(k6_relat_1(X1))&v1_funct_1(k6_relat_1(X1))), file('data/MPTP/proofs/747d1891215e47279b20aad99c2a9a2d.E_input', fc3_funct_1)).
fof(t71_relat_1, axiom, ![X1]:(k1_relat_1(k6_relat_1(X1))=X1&k2_relat_1(k6_relat_1(X1))=X1), file('data/MPTP/proofs/747d1891215e47279b20aad99c2a9a2d.E_input', t71_relat_1)).
fof(c_0_5, plain, ![X5, X6, X7]:(((r2_hidden(X5,k1_relat_1(X7))|~r2_hidden(k4_tarski(X5,X6),X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&(X6=k1_funct_1(X7,X5)|~r2_hidden(k4_tarski(X5,X6),X7)|(~v1_relat_1(X7)|~v1_funct_1(X7))))&(~r2_hidden(X5,k1_relat_1(X7))|X6!=k1_funct_1(X7,X5)|r2_hidden(k4_tarski(X5,X6),X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_funct_1])])])).
fof(c_0_6, plain, ![X8, X9, X10, X11, X12, X13]:((((r2_hidden(X10,X8)|~r2_hidden(k4_tarski(X10,X11),X9)|X9!=k6_relat_1(X8)|~v1_relat_1(X9))&(X10=X11|~r2_hidden(k4_tarski(X10,X11),X9)|X9!=k6_relat_1(X8)|~v1_relat_1(X9)))&(~r2_hidden(X12,X8)|X12!=X13|r2_hidden(k4_tarski(X12,X13),X9)|X9!=k6_relat_1(X8)|~v1_relat_1(X9)))&((~r2_hidden(k4_tarski(esk1_2(X8,X9),esk2_2(X8,X9)),X9)|(~r2_hidden(esk1_2(X8,X9),X8)|esk1_2(X8,X9)!=esk2_2(X8,X9))|X9=k6_relat_1(X8)|~v1_relat_1(X9))&((r2_hidden(esk1_2(X8,X9),X8)|r2_hidden(k4_tarski(esk1_2(X8,X9),esk2_2(X8,X9)),X9)|X9=k6_relat_1(X8)|~v1_relat_1(X9))&(esk1_2(X8,X9)=esk2_2(X8,X9)|r2_hidden(k4_tarski(esk1_2(X8,X9),esk2_2(X8,X9)),X9)|X9=k6_relat_1(X8)|~v1_relat_1(X9))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_relat_1])])])])])])).
fof(c_0_7, negated_conjecture, ~(![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(X2=k6_relat_1(X1)<=>(k1_relat_1(X2)=X1&![X3]:(r2_hidden(X3,X1)=>k1_funct_1(X2,X3)=X3))))), inference(assume_negation,[status(cth)],[t34_funct_1])).
cnf(c_0_8, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (esk1_2(X1,X2)=esk2_2(X1,X2)|r2_hidden(k4_tarski(esk1_2(X1,X2),esk2_2(X1,X2)),X2)|X2=k6_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_10, negated_conjecture, ![X19]:((v1_relat_1(esk4_0)&v1_funct_1(esk4_0))&(((r2_hidden(esk5_0,esk3_0)|k1_relat_1(esk4_0)!=esk3_0|esk4_0!=k6_relat_1(esk3_0))&(k1_funct_1(esk4_0,esk5_0)!=esk5_0|k1_relat_1(esk4_0)!=esk3_0|esk4_0!=k6_relat_1(esk3_0)))&((k1_relat_1(esk4_0)=esk3_0|esk4_0=k6_relat_1(esk3_0))&(~r2_hidden(X19,esk3_0)|k1_funct_1(esk4_0,X19)=X19|esk4_0=k6_relat_1(esk3_0))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).
cnf(c_0_11, plain, (X1=k1_funct_1(X2,X3)|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_12, plain, (esk2_2(X1,X2)=esk1_2(X1,X2)|X2=k6_relat_1(X1)|r2_hidden(esk1_2(X1,X2),k1_relat_1(X2))|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_13, negated_conjecture, (k1_relat_1(esk4_0)=esk3_0|esk4_0=k6_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (v1_relat_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (r2_hidden(k4_tarski(X1,X3),X2)|~r2_hidden(X1,k1_relat_1(X2))|X3!=k1_funct_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_17, plain, (r2_hidden(esk1_2(X1,X2),X1)|r2_hidden(k4_tarski(esk1_2(X1,X2),esk2_2(X1,X2)),X2)|X2=k6_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_18, negated_conjecture, (k1_funct_1(esk4_0,X1)=X1|esk4_0=k6_relat_1(esk3_0)|~r2_hidden(X1,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19, plain, (k1_funct_1(X1,esk1_2(X2,X1))=esk2_2(X2,X1)|esk2_2(X2,X1)=esk1_2(X2,X1)|X1=k6_relat_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_11, c_0_9])).
cnf(c_0_20, negated_conjecture, (esk2_2(X1,esk4_0)=esk1_2(X1,esk4_0)|k6_relat_1(esk3_0)=esk4_0|k6_relat_1(X1)=esk4_0|r2_hidden(esk1_2(X1,esk4_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_15])])).
cnf(c_0_21, plain, (r2_hidden(k4_tarski(X1,k1_funct_1(X2,X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (X1=k6_relat_1(X2)|r2_hidden(esk1_2(X2,X1),k1_relat_1(X1))|r2_hidden(esk1_2(X2,X1),X2)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_8, c_0_17])).
cnf(c_0_23, plain, (X2=k6_relat_1(X1)|~r2_hidden(k4_tarski(esk1_2(X1,X2),esk2_2(X1,X2)),X2)|~r2_hidden(esk1_2(X1,X2),X1)|esk1_2(X1,X2)!=esk2_2(X1,X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_24, negated_conjecture, (esk2_2(X1,esk4_0)=esk1_2(X1,esk4_0)|k6_relat_1(esk3_0)=esk4_0|k6_relat_1(X1)=esk4_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_14]), c_0_15])]), c_0_20])).
cnf(c_0_25, negated_conjecture, (k6_relat_1(esk3_0)=esk4_0|r2_hidden(k4_tarski(X1,X1),esk4_0)|~r2_hidden(X1,k1_relat_1(esk4_0))|~r2_hidden(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_18]), c_0_14]), c_0_15])])).
cnf(c_0_26, negated_conjecture, (k6_relat_1(X1)=esk4_0|r2_hidden(esk1_2(X1,esk4_0),k1_relat_1(esk4_0))|r2_hidden(esk1_2(X1,esk4_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_14]), c_0_15])])).
cnf(c_0_27, negated_conjecture, (k6_relat_1(esk3_0)=esk4_0|k6_relat_1(X1)=esk4_0|~r2_hidden(k4_tarski(esk1_2(X1,esk4_0),esk1_2(X1,esk4_0)),esk4_0)|~r2_hidden(esk1_2(X1,esk4_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_15])])).
cnf(c_0_28, negated_conjecture, (k6_relat_1(esk3_0)=esk4_0|r2_hidden(k4_tarski(X1,X1),esk4_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_25, c_0_13])).
cnf(c_0_29, negated_conjecture, (k6_relat_1(k1_relat_1(esk4_0))=esk4_0|r2_hidden(esk1_2(k1_relat_1(esk4_0),esk4_0),k1_relat_1(esk4_0))), inference(ef,[status(thm)],[c_0_26])).
fof(c_0_30, plain, ![X20]:(v1_relat_1(k6_relat_1(X20))&v1_funct_1(k6_relat_1(X20))), inference(variable_rename,[status(thm)],[fc3_funct_1])).
fof(c_0_31, plain, ![X21]:(k1_relat_1(k6_relat_1(X21))=X21&k2_relat_1(k6_relat_1(X21))=X21), inference(variable_rename,[status(thm)],[t71_relat_1])).
cnf(c_0_32, negated_conjecture, (k6_relat_1(esk3_0)=esk4_0|k6_relat_1(X1)=esk4_0|~r2_hidden(esk1_2(X1,esk4_0),esk3_0)|~r2_hidden(esk1_2(X1,esk4_0),X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_33, negated_conjecture, (k6_relat_1(esk3_0)=esk4_0|r2_hidden(esk1_2(esk3_0,esk4_0),esk3_0)), inference(spm,[status(thm)],[c_0_29, c_0_13])).
cnf(c_0_34, plain, (r2_hidden(k4_tarski(X1,X3),X4)|~r2_hidden(X1,X2)|X1!=X3|X4!=k6_relat_1(X2)|~v1_relat_1(X4)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_35, plain, (v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_36, plain, (k1_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_37, negated_conjecture, (k6_relat_1(esk3_0)=esk4_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_33])).
cnf(c_0_38, plain, (r2_hidden(k4_tarski(X1,X1),k6_relat_1(X2))|~r2_hidden(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_34])]), c_0_35])])).
cnf(c_0_39, plain, (v1_funct_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40, negated_conjecture, (k1_funct_1(esk4_0,esk5_0)!=esk5_0|k1_relat_1(esk4_0)!=esk3_0|esk4_0!=k6_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_41, negated_conjecture, (k1_relat_1(esk4_0)=esk3_0), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_42, plain, (k1_funct_1(k6_relat_1(X1),X2)=X2|~r2_hidden(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_38]), c_0_39]), c_0_35])])).
cnf(c_0_43, negated_conjecture, (r2_hidden(esk5_0,esk3_0)|k1_relat_1(esk4_0)!=esk3_0|esk4_0!=k6_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_44, negated_conjecture, (k1_funct_1(esk4_0,esk5_0)!=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_37])]), c_0_41])])).
cnf(c_0_45, negated_conjecture, (k1_funct_1(esk4_0,X1)=X1|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_42, c_0_37])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk5_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_37])]), c_0_41])])).
cnf(c_0_47, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 48
# Proof object clause steps            : 37
# Proof object formula steps           : 11
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 33
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 19
# Processed clauses                    : 82
# ...of these trivial                  : 0
# ...subsumed                          : 9
# ...remaining for further processing  : 73
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 18
# Generated clauses                    : 93
# ...of the previous two non-trivial   : 81
# Contextual simplify-reflections      : 2
# Paramodulations                      : 85
# Factorizations                       : 4
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 29
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 18
# Current number of unprocessed clauses: 35
# ...number of literals in the above   : 174
# Current number of archived formulas  : 0
# Current number of archived clauses   : 40
# Clause-clause subsumption calls (NU) : 479
# Rec. Clause-clause subsumption calls : 152
# Non-unit clause-clause subsumptions  : 14
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3356

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28708 pages
