# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/74dd22423bca4e32b464b94ea9bc1b35.E_input', t113_zfmisc_1)).
fof(d21_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X3=k15_lattice3(X1,X2)<=>(r4_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r4_lattice3(X1,X4,X2)=>r1_lattices(X1,X3,X4)))))))), file('data/MPTP/proofs/74dd22423bca4e32b464b94ea9bc1b35.E_input', d21_lattice3)).
fof(t152_zfmisc_1, axiom, ![X1, X2]:~(r2_hidden(X1,k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/74dd22423bca4e32b464b94ea9bc1b35.E_input', t152_zfmisc_1)).
fof(d16_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v13_lattices(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(X2=k5_lattices(X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_lattices(X1,X2,X3)=X2&k2_lattices(X1,X3,X2)=X2)))))), file('data/MPTP/proofs/74dd22423bca4e32b464b94ea9bc1b35.E_input', d16_lattices)).
fof(dt_k5_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>m1_subset_1(k5_lattices(X1),u1_struct_0(X1))), file('data/MPTP/proofs/74dd22423bca4e32b464b94ea9bc1b35.E_input', dt_k5_lattices)).
fof(t21_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattices(X1,X2,X3)<=>k2_lattices(X1,X2,X3)=X2)))), file('data/MPTP/proofs/74dd22423bca4e32b464b94ea9bc1b35.E_input', t21_lattices)).
fof(d18_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>(v4_lattice3(X1)<=>![X2]:?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r4_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r4_lattice3(X1,X4,X2)=>r1_lattices(X1,X3,X4))))))), file('data/MPTP/proofs/74dd22423bca4e32b464b94ea9bc1b35.E_input', d18_lattice3)).
fof(d17_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r4_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X4,X2)))))), file('data/MPTP/proofs/74dd22423bca4e32b464b94ea9bc1b35.E_input', d17_lattice3)).
fof(dt_l3_lattices, axiom, ![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1))), file('data/MPTP/proofs/74dd22423bca4e32b464b94ea9bc1b35.E_input', dt_l3_lattices)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/74dd22423bca4e32b464b94ea9bc1b35.E_input', cc1_lattices)).
fof(t50_lattice3, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))&k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0))), file('data/MPTP/proofs/74dd22423bca4e32b464b94ea9bc1b35.E_input', t50_lattice3)).
fof(d13_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v13_lattices(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_lattices(X1,X2,X3)=X2&k2_lattices(X1,X3,X2)=X2))))), file('data/MPTP/proofs/74dd22423bca4e32b464b94ea9bc1b35.E_input', d13_lattices)).
fof(d6_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v6_lattices(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2))))), file('data/MPTP/proofs/74dd22423bca4e32b464b94ea9bc1b35.E_input', d6_lattices)).
fof(c_0_13, plain, ![X11, X12]:((k2_zfmisc_1(X11,X12)!=k1_xboole_0|(X11=k1_xboole_0|X12=k1_xboole_0))&((X11!=k1_xboole_0|k2_zfmisc_1(X11,X12)=k1_xboole_0)&(X12!=k1_xboole_0|k2_zfmisc_1(X11,X12)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_14, plain, ![X6, X7, X8, X9]:(((r4_lattice3(X6,X8,X7)|X8!=k15_lattice3(X6,X7)|~m1_subset_1(X8,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6))))|(v2_struct_0(X6)|~l3_lattices(X6)))&(~m1_subset_1(X9,u1_struct_0(X6))|(~r4_lattice3(X6,X9,X7)|r1_lattices(X6,X8,X9))|X8!=k15_lattice3(X6,X7)|~m1_subset_1(X8,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6))))|(v2_struct_0(X6)|~l3_lattices(X6))))&((m1_subset_1(esk2_3(X6,X7,X8),u1_struct_0(X6))|~r4_lattice3(X6,X8,X7)|X8=k15_lattice3(X6,X7)|~m1_subset_1(X8,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6))))|(v2_struct_0(X6)|~l3_lattices(X6)))&((r4_lattice3(X6,esk2_3(X6,X7,X8),X7)|~r4_lattice3(X6,X8,X7)|X8=k15_lattice3(X6,X7)|~m1_subset_1(X8,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6))))|(v2_struct_0(X6)|~l3_lattices(X6)))&(~r1_lattices(X6,X8,esk2_3(X6,X7,X8))|~r4_lattice3(X6,X8,X7)|X8=k15_lattice3(X6,X7)|~m1_subset_1(X8,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6))))|(v2_struct_0(X6)|~l3_lattices(X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d21_lattice3])])])])])])).
fof(c_0_15, plain, ![X13, X14]:~r2_hidden(X13,k2_zfmisc_1(X13,X14)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_zfmisc_1])])).
cnf(c_0_16, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_17, plain, ![X38, X39, X40]:(((k2_lattices(X38,X39,X40)=X39|~m1_subset_1(X40,u1_struct_0(X38))|X39!=k5_lattices(X38)|~m1_subset_1(X39,u1_struct_0(X38))|~v13_lattices(X38)|(v2_struct_0(X38)|~l1_lattices(X38)))&(k2_lattices(X38,X40,X39)=X39|~m1_subset_1(X40,u1_struct_0(X38))|X39!=k5_lattices(X38)|~m1_subset_1(X39,u1_struct_0(X38))|~v13_lattices(X38)|(v2_struct_0(X38)|~l1_lattices(X38))))&((m1_subset_1(esk9_2(X38,X39),u1_struct_0(X38))|X39=k5_lattices(X38)|~m1_subset_1(X39,u1_struct_0(X38))|~v13_lattices(X38)|(v2_struct_0(X38)|~l1_lattices(X38)))&(k2_lattices(X38,X39,esk9_2(X38,X39))!=X39|k2_lattices(X38,esk9_2(X38,X39),X39)!=X39|X39=k5_lattices(X38)|~m1_subset_1(X39,u1_struct_0(X38))|~v13_lattices(X38)|(v2_struct_0(X38)|~l1_lattices(X38))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattices])])])])])])).
fof(c_0_18, plain, ![X37]:(v2_struct_0(X37)|~l1_lattices(X37)|m1_subset_1(k5_lattices(X37),u1_struct_0(X37))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_lattices])])])).
fof(c_0_19, plain, ![X20, X21, X22]:((~r1_lattices(X20,X21,X22)|k2_lattices(X20,X21,X22)=X21|~m1_subset_1(X22,u1_struct_0(X20))|~m1_subset_1(X21,u1_struct_0(X20))|(v2_struct_0(X20)|(~v8_lattices(X20)|(~v9_lattices(X20)|~l3_lattices(X20)))))&(k2_lattices(X20,X21,X22)!=X21|r1_lattices(X20,X21,X22)|~m1_subset_1(X22,u1_struct_0(X20))|~m1_subset_1(X21,u1_struct_0(X20))|(v2_struct_0(X20)|(~v8_lattices(X20)|(~v9_lattices(X20)|~l3_lattices(X20)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t21_lattices])])])])])).
fof(c_0_20, plain, ![X23, X24, X26, X28]:(((m1_subset_1(esk5_2(X23,X24),u1_struct_0(X23))|~v4_lattice3(X23)|(v2_struct_0(X23)|~l3_lattices(X23)))&((r4_lattice3(X23,esk5_2(X23,X24),X24)|~v4_lattice3(X23)|(v2_struct_0(X23)|~l3_lattices(X23)))&(~m1_subset_1(X26,u1_struct_0(X23))|(~r4_lattice3(X23,X26,X24)|r1_lattices(X23,esk5_2(X23,X24),X26))|~v4_lattice3(X23)|(v2_struct_0(X23)|~l3_lattices(X23)))))&((m1_subset_1(esk7_2(X23,X28),u1_struct_0(X23))|~r4_lattice3(X23,X28,esk6_1(X23))|~m1_subset_1(X28,u1_struct_0(X23))|v4_lattice3(X23)|(v2_struct_0(X23)|~l3_lattices(X23)))&((r4_lattice3(X23,esk7_2(X23,X28),esk6_1(X23))|~r4_lattice3(X23,X28,esk6_1(X23))|~m1_subset_1(X28,u1_struct_0(X23))|v4_lattice3(X23)|(v2_struct_0(X23)|~l3_lattices(X23)))&(~r1_lattices(X23,X28,esk7_2(X23,X28))|~r4_lattice3(X23,X28,esk6_1(X23))|~m1_subset_1(X28,u1_struct_0(X23))|v4_lattice3(X23)|(v2_struct_0(X23)|~l3_lattices(X23)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d18_lattice3])])])])])])).
cnf(c_0_21, plain, (X2=k15_lattice3(X1,X3)|v2_struct_0(X1)|v2_struct_0(X1)|~r1_lattices(X1,X2,esk2_3(X1,X3,X2))|~r4_lattice3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (~r2_hidden(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_16])).
fof(c_0_24, plain, ![X30, X31, X32, X33, X34]:((~r4_lattice3(X30,X31,X32)|(~m1_subset_1(X33,u1_struct_0(X30))|(~r2_hidden(X33,X32)|r1_lattices(X30,X33,X31)))|~m1_subset_1(X31,u1_struct_0(X30))|(v2_struct_0(X30)|~l3_lattices(X30)))&((m1_subset_1(esk8_3(X30,X31,X34),u1_struct_0(X30))|r4_lattice3(X30,X31,X34)|~m1_subset_1(X31,u1_struct_0(X30))|(v2_struct_0(X30)|~l3_lattices(X30)))&((r2_hidden(esk8_3(X30,X31,X34),X34)|r4_lattice3(X30,X31,X34)|~m1_subset_1(X31,u1_struct_0(X30))|(v2_struct_0(X30)|~l3_lattices(X30)))&(~r1_lattices(X30,esk8_3(X30,X31,X34),X31)|r4_lattice3(X30,X31,X34)|~m1_subset_1(X31,u1_struct_0(X30))|(v2_struct_0(X30)|~l3_lattices(X30)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_lattice3])])])])])])])).
cnf(c_0_25, plain, (k2_lattices(X1,X2,X3)=X3|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|X3!=k5_lattices(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v13_lattices(X1)|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (v2_struct_0(X1)|m1_subset_1(k5_lattices(X1),u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, plain, (k2_lattices(X1,X2,X3)=X2|v2_struct_0(X1)|~r1_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, plain, (r1_lattices(X2,esk5_2(X2,X3),X1)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r4_lattice3(X2,X1,X3)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_30, plain, ![X47]:((l1_lattices(X47)|~l3_lattices(X47))&(l2_lattices(X47)|~l3_lattices(X47))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
cnf(c_0_31, plain, (X2=k15_lattice3(X1,X3)|v2_struct_0(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~r4_lattice3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattices(X1,X2,esk2_3(X1,X3,X2))), inference(cn,[status(thm)],[c_0_21])).
cnf(c_0_32, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_33, plain, (r2_hidden(esk8_3(X1,X2,X3),X3)|r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, plain, (k2_lattices(X1,X2,k5_lattices(X1))=k5_lattices(X1)|v2_struct_0(X1)|~l1_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v13_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_25]), c_0_26])).
cnf(c_0_35, plain, (k2_lattices(X1,esk5_2(X1,X2),X3)=esk5_2(X1,X2)|v2_struct_0(X1)|~v9_lattices(X1)|~v8_lattices(X1)|~r4_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])).
cnf(c_0_36, plain, (l1_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37, plain, (r1_lattices(X2,X4,X1)|v2_struct_0(X2)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r4_lattice3(X2,X1,X3)|X4!=k15_lattice3(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_38, plain, (esk5_2(X1,X2)=k15_lattice3(X1,X3)|v2_struct_0(X1)|~r4_lattice3(X1,esk2_3(X1,X3,esk5_2(X1,X2)),X2)|~r4_lattice3(X1,esk5_2(X1,X2),X3)|~m1_subset_1(esk2_3(X1,X3,esk5_2(X1,X2)),u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_28]), c_0_29])).
cnf(c_0_39, plain, (r4_lattice3(X1,X2,k1_xboole_0)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_40, plain, (m1_subset_1(esk2_3(X1,X2,X3),u1_struct_0(X1))|X3=k15_lattice3(X1,X2)|v2_struct_0(X1)|v2_struct_0(X1)|~r4_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_41, plain, (esk5_2(X1,X2)=k5_lattices(X1)|v2_struct_0(X1)|~v9_lattices(X1)|~v8_lattices(X1)|~r4_lattice3(X1,k5_lattices(X1),X2)|~v13_lattices(X1)|~l3_lattices(X1)|~v4_lattice3(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_26]), c_0_29]), c_0_36])).
cnf(c_0_42, plain, (v2_struct_0(X2)|r1_lattices(X2,X4,X1)|X4!=k15_lattice3(X2,X3)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)|~r4_lattice3(X2,X1,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))), inference(cn,[status(thm)],[c_0_37])).
fof(c_0_43, plain, ![X36]:((~v2_struct_0(X36)|(v2_struct_0(X36)|~v10_lattices(X36))|~l3_lattices(X36))&((v4_lattices(X36)|(v2_struct_0(X36)|~v10_lattices(X36))|~l3_lattices(X36))&((v5_lattices(X36)|(v2_struct_0(X36)|~v10_lattices(X36))|~l3_lattices(X36))&((v6_lattices(X36)|(v2_struct_0(X36)|~v10_lattices(X36))|~l3_lattices(X36))&((v7_lattices(X36)|(v2_struct_0(X36)|~v10_lattices(X36))|~l3_lattices(X36))&((v8_lattices(X36)|(v2_struct_0(X36)|~v10_lattices(X36))|~l3_lattices(X36))&(v9_lattices(X36)|(v2_struct_0(X36)|~v10_lattices(X36))|~l3_lattices(X36)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
cnf(c_0_44, plain, (k15_lattice3(X1,X2)=esk5_2(X1,k1_xboole_0)|v2_struct_0(X1)|~r4_lattice3(X1,esk5_2(X1,k1_xboole_0),X2)|~m1_subset_1(esk2_3(X1,X2,esk5_2(X1,k1_xboole_0)),u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_45, plain, (X3=k15_lattice3(X1,X2)|v2_struct_0(X1)|m1_subset_1(esk2_3(X1,X2,X3),u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~r4_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_40])).
fof(c_0_46, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))&k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0)))), inference(assume_negation,[status(cth)],[t50_lattice3])).
cnf(c_0_47, plain, (esk5_2(X1,k1_xboole_0)=k5_lattices(X1)|v2_struct_0(X1)|~v9_lattices(X1)|~v8_lattices(X1)|~m1_subset_1(k5_lattices(X1),u1_struct_0(X1))|~v13_lattices(X1)|~l3_lattices(X1)|~v4_lattice3(X1)), inference(spm,[status(thm)],[c_0_41, c_0_39])).
fof(c_0_48, plain, ![X42, X44, X45]:(((m1_subset_1(esk10_1(X42),u1_struct_0(X42))|~v13_lattices(X42)|(v2_struct_0(X42)|~l1_lattices(X42)))&((k2_lattices(X42,esk10_1(X42),X44)=esk10_1(X42)|~m1_subset_1(X44,u1_struct_0(X42))|~v13_lattices(X42)|(v2_struct_0(X42)|~l1_lattices(X42)))&(k2_lattices(X42,X44,esk10_1(X42))=esk10_1(X42)|~m1_subset_1(X44,u1_struct_0(X42))|~v13_lattices(X42)|(v2_struct_0(X42)|~l1_lattices(X42)))))&((m1_subset_1(esk11_2(X42,X45),u1_struct_0(X42))|~m1_subset_1(X45,u1_struct_0(X42))|v13_lattices(X42)|(v2_struct_0(X42)|~l1_lattices(X42)))&(k2_lattices(X42,X45,esk11_2(X42,X45))!=X45|k2_lattices(X42,esk11_2(X42,X45),X45)!=X45|~m1_subset_1(X45,u1_struct_0(X42))|v13_lattices(X42)|(v2_struct_0(X42)|~l1_lattices(X42))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_lattices])])])])])])).
fof(c_0_49, plain, ![X15, X16, X17]:((~v6_lattices(X15)|(~m1_subset_1(X16,u1_struct_0(X15))|(~m1_subset_1(X17,u1_struct_0(X15))|k2_lattices(X15,X16,X17)=k2_lattices(X15,X17,X16)))|(v2_struct_0(X15)|~l1_lattices(X15)))&((m1_subset_1(esk3_1(X15),u1_struct_0(X15))|v6_lattices(X15)|(v2_struct_0(X15)|~l1_lattices(X15)))&((m1_subset_1(esk4_1(X15),u1_struct_0(X15))|v6_lattices(X15)|(v2_struct_0(X15)|~l1_lattices(X15)))&(k2_lattices(X15,esk3_1(X15),esk4_1(X15))!=k2_lattices(X15,esk4_1(X15),esk3_1(X15))|v6_lattices(X15)|(v2_struct_0(X15)|~l1_lattices(X15)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_lattices])])])])])])).
cnf(c_0_50, plain, (r1_lattices(X1,k15_lattice3(X1,X2),X3)|v2_struct_0(X1)|~r4_lattice3(X1,X3,X2)|~m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_51, plain, (v8_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_52, plain, (v9_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_53, plain, (k15_lattice3(X1,X2)=esk5_2(X1,k1_xboole_0)|v2_struct_0(X1)|~r4_lattice3(X1,esk5_2(X1,k1_xboole_0),X2)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_29])).
cnf(c_0_54, plain, (r4_lattice3(X1,esk5_2(X1,X2),X2)|v2_struct_0(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_55, negated_conjecture, ((~v2_struct_0(esk1_0)&(v10_lattices(esk1_0)&(v4_lattice3(esk1_0)&l3_lattices(esk1_0))))&(v2_struct_0(esk1_0)|(~v10_lattices(esk1_0)|(~v13_lattices(esk1_0)|~l3_lattices(esk1_0)))|k5_lattices(esk1_0)!=k15_lattice3(esk1_0,k1_xboole_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_46])])])])).
cnf(c_0_56, plain, (esk5_2(X1,k1_xboole_0)=k5_lattices(X1)|v2_struct_0(X1)|~v9_lattices(X1)|~v8_lattices(X1)|~v13_lattices(X1)|~l3_lattices(X1)|~v4_lattice3(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_26]), c_0_36])).
cnf(c_0_57, plain, (v13_lattices(X1)|v2_struct_0(X1)|k2_lattices(X1,X2,esk11_2(X1,X2))!=X2|k2_lattices(X1,esk11_2(X1,X2),X2)!=X2|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_58, plain, (k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2)|v2_struct_0(X1)|~v6_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_59, plain, (m1_subset_1(esk11_2(X1,X2),u1_struct_0(X1))|v13_lattices(X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_60, plain, (k2_lattices(X1,k15_lattice3(X1,X2),X3)=k15_lattice3(X1,X2)|v2_struct_0(X1)|~r4_lattice3(X1,X3,X2)|~m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_50]), c_0_51]), c_0_52])).
cnf(c_0_61, plain, (k15_lattice3(X1,k1_xboole_0)=esk5_2(X1,k1_xboole_0)|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_62, negated_conjecture, (v2_struct_0(esk1_0)|~v10_lattices(esk1_0)|~v13_lattices(esk1_0)|~l3_lattices(esk1_0)|k5_lattices(esk1_0)!=k15_lattice3(esk1_0,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_63, negated_conjecture, (v10_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_64, negated_conjecture, (l3_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_65, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_66, plain, (k15_lattice3(X1,X2)=k5_lattices(X1)|v2_struct_0(X1)|~r4_lattice3(X1,k5_lattices(X1),X2)|~v13_lattices(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_56]), c_0_51]), c_0_52])).
cnf(c_0_67, plain, (r4_lattice3(X1,k5_lattices(X1),k1_xboole_0)|v2_struct_0(X1)|~v9_lattices(X1)|~v8_lattices(X1)|~v13_lattices(X1)|~l3_lattices(X1)|~v4_lattice3(X1)), inference(spm,[status(thm)],[c_0_54, c_0_56])).
cnf(c_0_68, plain, (v13_lattices(X1)|v2_struct_0(X1)|k2_lattices(X1,X2,esk11_2(X1,X2))!=X2|~v6_lattices(X1)|~l1_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59])).
cnf(c_0_69, plain, (k2_lattices(X1,esk5_2(X1,k1_xboole_0),X2)=esk5_2(X1,k1_xboole_0)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_29]), c_0_39])).
cnf(c_0_70, plain, (v6_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_71, negated_conjecture, (k15_lattice3(esk1_0,k1_xboole_0)!=k5_lattices(esk1_0)|~v13_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_63]), c_0_64])]), c_0_65])).
cnf(c_0_72, plain, (k15_lattice3(X1,k1_xboole_0)=k5_lattices(X1)|v2_struct_0(X1)|~v13_lattices(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_51]), c_0_52])).
cnf(c_0_73, negated_conjecture, (v4_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_74, plain, (v13_lattices(X1)|v2_struct_0(X1)|~m1_subset_1(esk11_2(X1,esk5_2(X1,k1_xboole_0)),u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_29]), c_0_36]), c_0_70])).
cnf(c_0_75, negated_conjecture, (~v13_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_64]), c_0_73]), c_0_63])]), c_0_65])).
cnf(c_0_76, plain, (v13_lattices(X1)|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_59]), c_0_29]), c_0_36])).
cnf(c_0_77, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_64]), c_0_73]), c_0_63])]), c_0_65]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 78
# Proof object clause steps            : 51
# Proof object formula steps           : 27
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 42
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 49
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 48
# Processed clauses                    : 185
# ...of these trivial                  : 0
# ...subsumed                          : 72
# ...remaining for further processing  : 113
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 0
# Generated clauses                    : 246
# ...of the previous two non-trivial   : 215
# Contextual simplify-reflections      : 56
# Paramodulations                      : 240
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 104
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 95
# Current number of unprocessed clauses: 76
# ...number of literals in the above   : 781
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3
# Clause-clause subsumption calls (NU) : 5118
# Rec. Clause-clause subsumption calls : 191
# Non-unit clause-clause subsumptions  : 106
# Unit Clause-clause subsumption calls : 14
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 12263

# -------------------------------------------------
# User time                : 0.017 s
# System time              : 0.000 s
# Total time               : 0.017 s
# Maximum resident set size: 28876 pages
