# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(dt_k4_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&l1_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k4_lattices(X1,X2,X3),u1_struct_0(X1))), file('data/MPTP/proofs/755610ab515745b2bee559127257f958.E_input', dt_k4_lattices)).
fof(dt_l3_lattices, axiom, ![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1))), file('data/MPTP/proofs/755610ab515745b2bee559127257f958.E_input', dt_l3_lattices)).
fof(d8_filter_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v3_filter_0(X1)&l3_lattices(X1))))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k4_filter_0(X1,X2,X3)<=>(r3_lattices(X1,k4_lattices(X1,X2,X4),X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r3_lattices(X1,k4_lattices(X1,X2,X5),X3)=>r3_lattices(X1,X5,X4)))))))))), file('data/MPTP/proofs/755610ab515745b2bee559127257f958.E_input', d8_filter_0)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/755610ab515745b2bee559127257f958.E_input', cc1_lattices)).
fof(dt_k3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v4_lattices(X1)&l2_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k3_lattices(X1,X2,X3),u1_struct_0(X1))), file('data/MPTP/proofs/755610ab515745b2bee559127257f958.E_input', dt_k3_lattices)).
fof(commutativity_k4_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&l1_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k4_lattices(X1,X2,X3)=k4_lattices(X1,X3,X2)), file('data/MPTP/proofs/755610ab515745b2bee559127257f958.E_input', commutativity_k4_lattices)).
fof(t56_filter_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v3_filter_0(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_lattices(X1,k4_filter_0(X1,X2,X3),k4_filter_0(X1,X2,k3_lattices(X1,X3,X4)))&(r3_lattices(X1,k4_filter_0(X1,X2,X3),k4_filter_0(X1,k4_lattices(X1,X2,X4),X3))&(r3_lattices(X1,k4_filter_0(X1,X2,X3),k4_filter_0(X1,X2,k3_lattices(X1,X4,X3)))&r3_lattices(X1,k4_filter_0(X1,X2,X3),k4_filter_0(X1,k4_lattices(X1,X4,X2),X3))))))))), file('data/MPTP/proofs/755610ab515745b2bee559127257f958.E_input', t56_filter_1)).
fof(t3_filter_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_lattices(X1,X2,X3)=>r3_lattices(X1,X2,k3_lattices(X1,X4,X3))))))), file('data/MPTP/proofs/755610ab515745b2bee559127257f958.E_input', t3_filter_0)).
fof(d3_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l2_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattices(X1,X2,X3)<=>k1_lattices(X1,X2,X3)=X3)))), file('data/MPTP/proofs/755610ab515745b2bee559127257f958.E_input', d3_lattices)).
fof(t23_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>r1_lattices(X1,k4_lattices(X1,X2,X3),X2)))), file('data/MPTP/proofs/755610ab515745b2bee559127257f958.E_input', t23_lattices)).
fof(reflexivity_r3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>r3_lattices(X1,X2,X2)), file('data/MPTP/proofs/755610ab515745b2bee559127257f958.E_input', reflexivity_r3_lattices)).
fof(commutativity_k3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v4_lattices(X1)&l2_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k3_lattices(X1,X2,X3)=k3_lattices(X1,X3,X2)), file('data/MPTP/proofs/755610ab515745b2bee559127257f958.E_input', commutativity_k3_lattices)).
fof(redefinition_k3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v4_lattices(X1)&l2_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k3_lattices(X1,X2,X3)=k1_lattices(X1,X2,X3)), file('data/MPTP/proofs/755610ab515745b2bee559127257f958.E_input', redefinition_k3_lattices)).
fof(dt_k4_filter_0, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k4_filter_0(X1,X2,X3),u1_struct_0(X1))), file('data/MPTP/proofs/755610ab515745b2bee559127257f958.E_input', dt_k4_filter_0)).
fof(t17_filter_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v3_filter_0(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>k4_filter_0(X1,k4_lattices(X1,X2,X3),X4)=k4_filter_0(X1,X2,k4_filter_0(X1,X3,X4)))))), file('data/MPTP/proofs/755610ab515745b2bee559127257f958.E_input', t17_filter_1)).
fof(c_0_15, plain, ![X10, X11, X12]:(v2_struct_0(X10)|(~v6_lattices(X10)|~l1_lattices(X10))|(~m1_subset_1(X11,u1_struct_0(X10))|~m1_subset_1(X12,u1_struct_0(X10)))|m1_subset_1(k4_lattices(X10,X11,X12),u1_struct_0(X10))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_lattices])])])).
fof(c_0_16, plain, ![X26]:((l1_lattices(X26)|~l3_lattices(X26))&(l2_lattices(X26)|~l3_lattices(X26))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
fof(c_0_17, plain, ![X16, X17, X18, X19, X20]:(((r3_lattices(X16,k4_lattices(X16,X17,X19),X18)|X19!=k4_filter_0(X16,X17,X18)|~m1_subset_1(X19,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v3_filter_0(X16)|~l3_lattices(X16))))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|~l3_lattices(X16))))&(~m1_subset_1(X20,u1_struct_0(X16))|(~r3_lattices(X16,k4_lattices(X16,X17,X20),X18)|r3_lattices(X16,X20,X19))|X19!=k4_filter_0(X16,X17,X18)|~m1_subset_1(X19,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v3_filter_0(X16)|~l3_lattices(X16))))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|~l3_lattices(X16)))))&((m1_subset_1(esk5_4(X16,X17,X18,X19),u1_struct_0(X16))|~r3_lattices(X16,k4_lattices(X16,X17,X19),X18)|X19=k4_filter_0(X16,X17,X18)|~m1_subset_1(X19,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v3_filter_0(X16)|~l3_lattices(X16))))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|~l3_lattices(X16))))&((r3_lattices(X16,k4_lattices(X16,X17,esk5_4(X16,X17,X18,X19)),X18)|~r3_lattices(X16,k4_lattices(X16,X17,X19),X18)|X19=k4_filter_0(X16,X17,X18)|~m1_subset_1(X19,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v3_filter_0(X16)|~l3_lattices(X16))))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|~l3_lattices(X16))))&(~r3_lattices(X16,esk5_4(X16,X17,X18,X19),X19)|~r3_lattices(X16,k4_lattices(X16,X17,X19),X18)|X19=k4_filter_0(X16,X17,X18)|~m1_subset_1(X19,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v3_filter_0(X16)|~l3_lattices(X16))))|~m1_subset_1(X18,u1_struct_0(X16))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|~l3_lattices(X16))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_filter_0])])])])])])).
cnf(c_0_18, plain, (v2_struct_0(X1)|m1_subset_1(k4_lattices(X1,X2,X3),u1_struct_0(X1))|~v6_lattices(X1)|~l1_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, plain, (l1_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_20, plain, ![X34]:((~v2_struct_0(X34)|(v2_struct_0(X34)|~v10_lattices(X34))|~l3_lattices(X34))&((v4_lattices(X34)|(v2_struct_0(X34)|~v10_lattices(X34))|~l3_lattices(X34))&((v5_lattices(X34)|(v2_struct_0(X34)|~v10_lattices(X34))|~l3_lattices(X34))&((v6_lattices(X34)|(v2_struct_0(X34)|~v10_lattices(X34))|~l3_lattices(X34))&((v7_lattices(X34)|(v2_struct_0(X34)|~v10_lattices(X34))|~l3_lattices(X34))&((v8_lattices(X34)|(v2_struct_0(X34)|~v10_lattices(X34))|~l3_lattices(X34))&(v9_lattices(X34)|(v2_struct_0(X34)|~v10_lattices(X34))|~l3_lattices(X34)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
fof(c_0_21, plain, ![X47, X48, X49]:(v2_struct_0(X47)|(~v4_lattices(X47)|~l2_lattices(X47))|(~m1_subset_1(X48,u1_struct_0(X47))|~m1_subset_1(X49,u1_struct_0(X47)))|m1_subset_1(k3_lattices(X47,X48,X49),u1_struct_0(X47))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_lattices])])])).
fof(c_0_22, plain, ![X41, X42, X43]:(v2_struct_0(X41)|(~v6_lattices(X41)|~l1_lattices(X41))|(~m1_subset_1(X42,u1_struct_0(X41))|~m1_subset_1(X43,u1_struct_0(X41)))|k4_lattices(X41,X42,X43)=k4_lattices(X41,X43,X42)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[commutativity_k4_lattices])])])).
fof(c_0_23, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v3_filter_0(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_lattices(X1,k4_filter_0(X1,X2,X3),k4_filter_0(X1,X2,k3_lattices(X1,X3,X4)))&(r3_lattices(X1,k4_filter_0(X1,X2,X3),k4_filter_0(X1,k4_lattices(X1,X2,X4),X3))&(r3_lattices(X1,k4_filter_0(X1,X2,X3),k4_filter_0(X1,X2,k3_lattices(X1,X4,X3)))&r3_lattices(X1,k4_filter_0(X1,X2,X3),k4_filter_0(X1,k4_lattices(X1,X4,X2),X3)))))))))), inference(assume_negation,[status(cth)],[t56_filter_1])).
cnf(c_0_24, plain, (r3_lattices(X2,X1,X5)|v2_struct_0(X2)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r3_lattices(X2,k4_lattices(X2,X3,X1),X4)|X5!=k4_filter_0(X2,X3,X4)|~m1_subset_1(X5,u1_struct_0(X2))|~v10_lattices(X2)|~v3_filter_0(X2)|~l3_lattices(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~v10_lattices(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_25, plain, ![X30, X31, X32, X33]:(v2_struct_0(X30)|(~v10_lattices(X30)|~l3_lattices(X30))|(~m1_subset_1(X31,u1_struct_0(X30))|(~m1_subset_1(X32,u1_struct_0(X30))|(~m1_subset_1(X33,u1_struct_0(X30))|(~r3_lattices(X30,X31,X32)|r3_lattices(X30,X31,k3_lattices(X30,X33,X32))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_filter_0])])])])).
cnf(c_0_26, plain, (m1_subset_1(k4_lattices(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~v6_lattices(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_27, plain, (v6_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (v2_struct_0(X1)|m1_subset_1(k3_lattices(X1,X2,X3),u1_struct_0(X1))|~v4_lattices(X1)|~l2_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, plain, (v4_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30, plain, (l2_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_31, plain, ![X35, X36, X37]:((~r1_lattices(X35,X36,X37)|k1_lattices(X35,X36,X37)=X37|~m1_subset_1(X37,u1_struct_0(X35))|~m1_subset_1(X36,u1_struct_0(X35))|(v2_struct_0(X35)|~l2_lattices(X35)))&(k1_lattices(X35,X36,X37)!=X37|r1_lattices(X35,X36,X37)|~m1_subset_1(X37,u1_struct_0(X35))|~m1_subset_1(X36,u1_struct_0(X35))|(v2_struct_0(X35)|~l2_lattices(X35)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d3_lattices])])])])])).
fof(c_0_32, plain, ![X44, X45, X46]:(v2_struct_0(X44)|(~v6_lattices(X44)|(~v8_lattices(X44)|~l3_lattices(X44)))|(~m1_subset_1(X45,u1_struct_0(X44))|(~m1_subset_1(X46,u1_struct_0(X44))|r1_lattices(X44,k4_lattices(X44,X45,X46),X45)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t23_lattices])])])])).
cnf(c_0_33, plain, (v2_struct_0(X1)|k4_lattices(X1,X2,X3)=k4_lattices(X1,X3,X2)|~v6_lattices(X1)|~l1_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_34, plain, ![X13, X14, X15]:(v2_struct_0(X13)|(~v6_lattices(X13)|(~v8_lattices(X13)|(~v9_lattices(X13)|~l3_lattices(X13))))|(~m1_subset_1(X14,u1_struct_0(X13))|~m1_subset_1(X15,u1_struct_0(X13)))|r3_lattices(X13,X14,X14)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r3_lattices])])])).
fof(c_0_35, negated_conjecture, ((~v2_struct_0(esk1_0)&(v10_lattices(esk1_0)&(v3_filter_0(esk1_0)&l3_lattices(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(m1_subset_1(esk4_0,u1_struct_0(esk1_0))&(~r3_lattices(esk1_0,k4_filter_0(esk1_0,esk2_0,esk3_0),k4_filter_0(esk1_0,esk2_0,k3_lattices(esk1_0,esk3_0,esk4_0)))|(~r3_lattices(esk1_0,k4_filter_0(esk1_0,esk2_0,esk3_0),k4_filter_0(esk1_0,k4_lattices(esk1_0,esk2_0,esk4_0),esk3_0))|(~r3_lattices(esk1_0,k4_filter_0(esk1_0,esk2_0,esk3_0),k4_filter_0(esk1_0,esk2_0,k3_lattices(esk1_0,esk4_0,esk3_0)))|~r3_lattices(esk1_0,k4_filter_0(esk1_0,esk2_0,esk3_0),k4_filter_0(esk1_0,k4_lattices(esk1_0,esk4_0,esk2_0),esk3_0))))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_23])])])])).
cnf(c_0_36, plain, (v2_struct_0(X2)|r3_lattices(X2,X1,X5)|X5!=k4_filter_0(X2,X3,X4)|~v10_lattices(X2)|~v3_filter_0(X2)|~l3_lattices(X2)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~r3_lattices(X2,k4_lattices(X2,X3,X1),X4)), inference(cn,[status(thm)],[c_0_24])).
cnf(c_0_37, plain, (v2_struct_0(X1)|r3_lattices(X1,X2,k3_lattices(X1,X4,X3))|~v10_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~r3_lattices(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_38, plain, (m1_subset_1(k4_lattices(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_39, plain, (m1_subset_1(k3_lattices(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])).
cnf(c_0_40, plain, (r3_lattices(X1,k4_lattices(X1,X2,X3),X4)|v2_struct_0(X1)|v2_struct_0(X1)|X3!=k4_filter_0(X1,X2,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~v10_lattices(X1)|~v3_filter_0(X1)|~l3_lattices(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_41, plain, ![X38, X39, X40]:(v2_struct_0(X38)|(~v4_lattices(X38)|~l2_lattices(X38))|(~m1_subset_1(X39,u1_struct_0(X38))|~m1_subset_1(X40,u1_struct_0(X38)))|k3_lattices(X38,X39,X40)=k3_lattices(X38,X40,X39)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[commutativity_k3_lattices])])])).
cnf(c_0_42, plain, (k1_lattices(X1,X2,X3)=X3|v2_struct_0(X1)|~r1_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_43, plain, (v2_struct_0(X1)|r1_lattices(X1,k4_lattices(X1,X2,X3),X2)|~v6_lattices(X1)|~v8_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_44, plain, (k4_lattices(X1,X2,X3)=k4_lattices(X1,X3,X2)|v2_struct_0(X1)|~v6_lattices(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_33, c_0_19])).
cnf(c_0_45, plain, (v2_struct_0(X1)|r3_lattices(X1,X2,X2)|~v6_lattices(X1)|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_47, negated_conjecture, (l3_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_48, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_49, plain, (r3_lattices(X1,X2,X3)|v2_struct_0(X1)|X3!=k4_filter_0(X1,X4,k3_lattices(X1,X5,X6))|~r3_lattices(X1,k4_lattices(X1,X4,X2),X6)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~l3_lattices(X1)|~v3_filter_0(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_39])).
cnf(c_0_50, plain, (v2_struct_0(X1)|r3_lattices(X1,k4_lattices(X1,X2,X3),X4)|X3!=k4_filter_0(X1,X2,X4)|~v10_lattices(X1)|~v3_filter_0(X1)|~l3_lattices(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_40])).
cnf(c_0_51, negated_conjecture, (v3_filter_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_52, negated_conjecture, (v10_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_53, plain, (v2_struct_0(X1)|k3_lattices(X1,X2,X3)=k3_lattices(X1,X3,X2)|~v4_lattices(X1)|~l2_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_54, plain, ![X50, X51, X52]:(v2_struct_0(X50)|(~v4_lattices(X50)|~l2_lattices(X50))|(~m1_subset_1(X51,u1_struct_0(X50))|~m1_subset_1(X52,u1_struct_0(X50)))|k3_lattices(X50,X51,X52)=k1_lattices(X50,X51,X52)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_lattices])])])).
cnf(c_0_55, plain, (k1_lattices(X1,k4_lattices(X1,X2,X3),X2)=X2|v2_struct_0(X1)|~v8_lattices(X1)|~v6_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_26]), c_0_30])).
cnf(c_0_56, plain, (k4_lattices(X1,X2,X3)=k4_lattices(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_44, c_0_27])).
cnf(c_0_57, plain, (v8_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_58, negated_conjecture, (r3_lattices(esk1_0,X1,X1)|~v9_lattices(esk1_0)|~v8_lattices(esk1_0)|~v6_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])]), c_0_48])).
cnf(c_0_59, plain, (v9_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_60, plain, (r3_lattices(X1,X2,k4_filter_0(X1,X3,k3_lattices(X1,X4,X5)))|v2_struct_0(X1)|~r3_lattices(X1,k4_lattices(X1,X3,X2),X5)|~m1_subset_1(k4_filter_0(X1,X3,k3_lattices(X1,X4,X5)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~l3_lattices(X1)|~v3_filter_0(X1)|~v10_lattices(X1)), inference(er,[status(thm)],[c_0_49])).
cnf(c_0_61, negated_conjecture, (r3_lattices(esk1_0,k4_lattices(esk1_0,X1,X2),X3)|X2!=k4_filter_0(esk1_0,X1,X3)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_47]), c_0_52])]), c_0_48])).
cnf(c_0_62, plain, (k3_lattices(X1,X2,X3)=k3_lattices(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_29]), c_0_30])).
cnf(c_0_63, plain, (v2_struct_0(X1)|k3_lattices(X1,X2,X3)=k1_lattices(X1,X2,X3)|~v4_lattices(X1)|~l2_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_64, plain, (k1_lattices(X1,k4_lattices(X1,X2,X3),X3)=X3|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_27]), c_0_57])).
cnf(c_0_65, negated_conjecture, (r3_lattices(esk1_0,X1,X1)|~v8_lattices(esk1_0)|~v6_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_47]), c_0_52])]), c_0_48])).
cnf(c_0_66, negated_conjecture, (~r3_lattices(esk1_0,k4_filter_0(esk1_0,esk2_0,esk3_0),k4_filter_0(esk1_0,esk2_0,k3_lattices(esk1_0,esk3_0,esk4_0)))|~r3_lattices(esk1_0,k4_filter_0(esk1_0,esk2_0,esk3_0),k4_filter_0(esk1_0,k4_lattices(esk1_0,esk2_0,esk4_0),esk3_0))|~r3_lattices(esk1_0,k4_filter_0(esk1_0,esk2_0,esk3_0),k4_filter_0(esk1_0,esk2_0,k3_lattices(esk1_0,esk4_0,esk3_0)))|~r3_lattices(esk1_0,k4_filter_0(esk1_0,esk2_0,esk3_0),k4_filter_0(esk1_0,k4_lattices(esk1_0,esk4_0,esk2_0),esk3_0))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_67, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_68, negated_conjecture, (r3_lattices(esk1_0,X1,k4_filter_0(esk1_0,X2,k3_lattices(esk1_0,X3,X4)))|X1!=k4_filter_0(esk1_0,X2,X4)|~m1_subset_1(k4_filter_0(esk1_0,X2,k3_lattices(esk1_0,X3,X4)),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_47]), c_0_51]), c_0_52])]), c_0_48])).
fof(c_0_69, plain, ![X27, X28, X29]:(v2_struct_0(X27)|(~v10_lattices(X27)|~l3_lattices(X27))|(~m1_subset_1(X28,u1_struct_0(X27))|~m1_subset_1(X29,u1_struct_0(X27)))|m1_subset_1(k4_filter_0(X27,X28,X29),u1_struct_0(X27))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_filter_0])])])).
cnf(c_0_70, plain, (r3_lattices(X1,X2,k3_lattices(X1,X3,X4))|v2_struct_0(X1)|~r3_lattices(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_37, c_0_62])).
cnf(c_0_71, plain, (k3_lattices(X1,k4_lattices(X1,X2,X3),X3)=X3|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_38]), c_0_30]), c_0_29])).
cnf(c_0_72, negated_conjecture, (r3_lattices(esk1_0,X1,X1)|~v6_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_57]), c_0_47]), c_0_52])]), c_0_48])).
cnf(c_0_73, negated_conjecture, (~r3_lattices(esk1_0,k4_filter_0(esk1_0,esk2_0,esk3_0),k4_filter_0(esk1_0,esk2_0,k3_lattices(esk1_0,esk3_0,esk4_0)))|~r3_lattices(esk1_0,k4_filter_0(esk1_0,esk2_0,esk3_0),k4_filter_0(esk1_0,k4_lattices(esk1_0,esk2_0,esk4_0),esk3_0))|~r3_lattices(esk1_0,k4_filter_0(esk1_0,esk2_0,esk3_0),k4_filter_0(esk1_0,k4_lattices(esk1_0,esk4_0,esk2_0),esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_62]), c_0_46]), c_0_67]), c_0_47]), c_0_52])]), c_0_48])).
cnf(c_0_74, negated_conjecture, (r3_lattices(esk1_0,X1,k4_filter_0(esk1_0,X2,k3_lattices(esk1_0,X3,X4)))|X1!=k4_filter_0(esk1_0,X2,X3)|~m1_subset_1(k4_filter_0(esk1_0,X2,k3_lattices(esk1_0,X3,X4)),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_62]), c_0_47]), c_0_52])]), c_0_48])).
cnf(c_0_75, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_76, plain, (v2_struct_0(X1)|m1_subset_1(k4_filter_0(X1,X2,X3),u1_struct_0(X1))|~v10_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_77, plain, (r3_lattices(X1,X2,X3)|v2_struct_0(X1)|~r3_lattices(X1,X2,k4_lattices(X1,X4,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_38])).
cnf(c_0_78, negated_conjecture, (r3_lattices(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_27]), c_0_47]), c_0_52])]), c_0_48])).
fof(c_0_79, plain, ![X22, X23, X24, X25]:(v2_struct_0(X22)|(~v10_lattices(X22)|(~v3_filter_0(X22)|~l3_lattices(X22)))|(~m1_subset_1(X23,u1_struct_0(X22))|(~m1_subset_1(X24,u1_struct_0(X22))|(~m1_subset_1(X25,u1_struct_0(X22))|k4_filter_0(X22,k4_lattices(X22,X23,X24),X25)=k4_filter_0(X22,X23,k4_filter_0(X22,X24,X25)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_filter_1])])])])).
cnf(c_0_80, negated_conjecture, (~r3_lattices(esk1_0,k4_filter_0(esk1_0,esk2_0,esk3_0),k4_filter_0(esk1_0,k4_lattices(esk1_0,esk2_0,esk4_0),esk3_0))|~r3_lattices(esk1_0,k4_filter_0(esk1_0,esk2_0,esk3_0),k4_filter_0(esk1_0,k4_lattices(esk1_0,esk4_0,esk2_0),esk3_0))|~m1_subset_1(k4_filter_0(esk1_0,esk2_0,k3_lattices(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk1_0))|~m1_subset_1(k4_filter_0(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_75]), c_0_46]), c_0_67])])).
cnf(c_0_81, negated_conjecture, (m1_subset_1(k4_filter_0(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_52]), c_0_47])]), c_0_48])).
cnf(c_0_82, negated_conjecture, (r3_lattices(esk1_0,k4_lattices(esk1_0,X1,X2),X2)|~m1_subset_1(k4_lattices(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_47]), c_0_52])]), c_0_48])).
cnf(c_0_83, plain, (v2_struct_0(X1)|k4_filter_0(X1,k4_lattices(X1,X2,X3),X4)=k4_filter_0(X1,X2,k4_filter_0(X1,X3,X4))|~v10_lattices(X1)|~v3_filter_0(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_84, negated_conjecture, (~r3_lattices(esk1_0,k4_filter_0(esk1_0,esk2_0,esk3_0),k4_filter_0(esk1_0,k4_lattices(esk1_0,esk2_0,esk4_0),esk3_0))|~r3_lattices(esk1_0,k4_filter_0(esk1_0,esk2_0,esk3_0),k4_filter_0(esk1_0,k4_lattices(esk1_0,esk4_0,esk2_0),esk3_0))|~m1_subset_1(k4_filter_0(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_lattices(esk1_0,esk3_0,esk4_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_75])])).
cnf(c_0_85, negated_conjecture, (r3_lattices(esk1_0,X1,X2)|X2!=k4_filter_0(esk1_0,X3,X1)|~m1_subset_1(k4_lattices(esk1_0,X3,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_82]), c_0_47]), c_0_51]), c_0_52])]), c_0_48])).
cnf(c_0_86, negated_conjecture, (k4_filter_0(esk1_0,k4_lattices(esk1_0,X1,X2),X3)=k4_filter_0(esk1_0,X1,k4_filter_0(esk1_0,X2,X3))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_51]), c_0_47]), c_0_52])]), c_0_48])).
cnf(c_0_87, negated_conjecture, (~r3_lattices(esk1_0,k4_filter_0(esk1_0,esk2_0,esk3_0),k4_filter_0(esk1_0,k4_lattices(esk1_0,esk2_0,esk4_0),esk3_0))|~r3_lattices(esk1_0,k4_filter_0(esk1_0,esk2_0,esk3_0),k4_filter_0(esk1_0,k4_lattices(esk1_0,esk4_0,esk2_0),esk3_0))|~m1_subset_1(k4_filter_0(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_39]), c_0_46]), c_0_67]), c_0_47]), c_0_52])]), c_0_48])).
cnf(c_0_88, negated_conjecture, (r3_lattices(esk1_0,X1,X2)|X2!=k4_filter_0(esk1_0,X3,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_38]), c_0_47]), c_0_52])]), c_0_48])).
cnf(c_0_89, negated_conjecture, (k4_filter_0(esk1_0,k4_lattices(esk1_0,X1,X2),X3)=k4_filter_0(esk1_0,X2,k4_filter_0(esk1_0,X1,X3))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_56]), c_0_47]), c_0_52])]), c_0_48])).
cnf(c_0_90, negated_conjecture, (~r3_lattices(esk1_0,k4_filter_0(esk1_0,esk2_0,esk3_0),k4_filter_0(esk1_0,k4_lattices(esk1_0,esk2_0,esk4_0),esk3_0))|~m1_subset_1(k4_filter_0(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_56]), c_0_75]), c_0_46]), c_0_47]), c_0_52])]), c_0_48])).
cnf(c_0_91, negated_conjecture, (r3_lattices(esk1_0,X1,k4_filter_0(esk1_0,X2,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_88]), c_0_81])).
cnf(c_0_92, negated_conjecture, (k4_filter_0(esk1_0,X1,k4_filter_0(esk1_0,X2,X3))=k4_filter_0(esk1_0,X2,k4_filter_0(esk1_0,X1,X3))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_86, c_0_89])).
cnf(c_0_93, negated_conjecture, (~r3_lattices(esk1_0,k4_filter_0(esk1_0,esk2_0,esk3_0),k4_filter_0(esk1_0,esk2_0,k4_filter_0(esk1_0,esk4_0,esk3_0)))|~m1_subset_1(k4_filter_0(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_86]), c_0_67]), c_0_46]), c_0_75])])).
cnf(c_0_94, negated_conjecture, (r3_lattices(esk1_0,k4_filter_0(esk1_0,X1,X2),k4_filter_0(esk1_0,X1,k4_filter_0(esk1_0,X3,X2)))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_81])).
cnf(c_0_95, negated_conjecture, (~m1_subset_1(k4_filter_0(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_46]), c_0_67]), c_0_75])])).
cnf(c_0_96, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_81]), c_0_67]), c_0_75])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 97
# Proof object clause steps            : 66
# Proof object formula steps           : 31
# Proof object conjectures             : 35
# Proof object clause conjectures      : 32
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 37
# Proof object simplifying inferences  : 103
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 33
# Processed clauses                    : 1729
# ...of these trivial                  : 0
# ...subsumed                          : 1272
# ...remaining for further processing  : 457
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 101
# Backward-rewritten                   : 0
# Generated clauses                    : 14243
# ...of the previous two non-trivial   : 14232
# Contextual simplify-reflections      : 99
# Paramodulations                      : 14099
# Factorizations                       : 0
# Equation resolutions                 : 144
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 323
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 315
# Current number of unprocessed clauses: 12419
# ...number of literals in the above   : 125899
# Current number of archived formulas  : 0
# Current number of archived clauses   : 134
# Clause-clause subsumption calls (NU) : 79688
# Rec. Clause-clause subsumption calls : 19943
# Non-unit clause-clause subsumptions  : 1472
# Unit Clause-clause subsumption calls : 196
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 702835

# -------------------------------------------------
# User time                : 0.291 s
# System time              : 0.003 s
# Total time               : 0.294 s
# Maximum resident set size: 28688 pages
