# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_S0e
# and selection function SelectLargestNegLit.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t56_tex_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v3_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_xboole_0(k2_pre_topc(X1,k6_domain_1(u1_struct_0(X1),X2)),k2_pre_topc(X1,k6_domain_1(u1_struct_0(X1),X3)))|k2_pre_topc(X1,k6_domain_1(u1_struct_0(X1),X2))=k2_pre_topc(X1,k6_domain_1(u1_struct_0(X1),X3)))))), file('data/MPTP/proofs/764abd797b1a48f68d08f616dafab43b.E_input', t56_tex_2)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/764abd797b1a48f68d08f616dafab43b.E_input', t2_subset)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/764abd797b1a48f68d08f616dafab43b.E_input', dt_l1_pre_topc)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/764abd797b1a48f68d08f616dafab43b.E_input', fc2_struct_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/764abd797b1a48f68d08f616dafab43b.E_input', d1_xboole_0)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/764abd797b1a48f68d08f616dafab43b.E_input', redefinition_k6_domain_1)).
fof(t63_subset_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))), file('data/MPTP/proofs/764abd797b1a48f68d08f616dafab43b.E_input', t63_subset_1)).
fof(dt_k2_pre_topc, axiom, ![X1, X2]:((l1_pre_topc(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k2_pre_topc(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/764abd797b1a48f68d08f616dafab43b.E_input', dt_k2_pre_topc)).
fof(fc1_tops_1, axiom, ![X1, X2]:(((v2_pre_topc(X1)&l1_pre_topc(X1))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>v4_pre_topc(k2_pre_topc(X1,X2),X1)), file('data/MPTP/proofs/764abd797b1a48f68d08f616dafab43b.E_input', fc1_tops_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/764abd797b1a48f68d08f616dafab43b.E_input', d1_tarski)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/764abd797b1a48f68d08f616dafab43b.E_input', t3_xboole_0)).
fof(t40_tsep_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r1_xboole_0(X2,X3)&v3_pre_topc(X2,X1))=>r1_xboole_0(X2,k2_pre_topc(X1,X3)))))), file('data/MPTP/proofs/764abd797b1a48f68d08f616dafab43b.E_input', t40_tsep_1)).
fof(t24_tdlat_3, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>(v3_tdlat_3(X1)<=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v4_pre_topc(X2,X1)=>v3_pre_topc(X2,X1))))), file('data/MPTP/proofs/764abd797b1a48f68d08f616dafab43b.E_input', t24_tdlat_3)).
fof(t55_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v3_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X2,k2_pre_topc(X1,k6_domain_1(u1_struct_0(X1),X3)))=>k2_pre_topc(X1,k6_domain_1(u1_struct_0(X1),X2))=k2_pre_topc(X1,k6_domain_1(u1_struct_0(X1),X3)))))), file('data/MPTP/proofs/764abd797b1a48f68d08f616dafab43b.E_input', t55_tex_2)).
fof(c_0_14, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v3_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_xboole_0(k2_pre_topc(X1,k6_domain_1(u1_struct_0(X1),X2)),k2_pre_topc(X1,k6_domain_1(u1_struct_0(X1),X3)))|k2_pre_topc(X1,k6_domain_1(u1_struct_0(X1),X2))=k2_pre_topc(X1,k6_domain_1(u1_struct_0(X1),X3))))))), inference(assume_negation,[status(cth)],[t56_tex_2])).
fof(c_0_15, plain, ![X7, X8]:(~m1_subset_1(X7,X8)|(v1_xboole_0(X8)|r2_hidden(X7,X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_16, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&(v3_tdlat_3(esk1_0)&l1_pre_topc(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(~r1_xboole_0(k2_pre_topc(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0)),k2_pre_topc(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))&k2_pre_topc(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=k2_pre_topc(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])).
fof(c_0_17, plain, ![X9]:(~l1_pre_topc(X9)|l1_struct_0(X9)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_18, plain, ![X10]:(v2_struct_0(X10)|~l1_struct_0(X10)|~v1_xboole_0(u1_struct_0(X10))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_19, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_23, plain, ![X11, X12, X13]:((~v1_xboole_0(X11)|~r2_hidden(X12,X11))&(r2_hidden(esk4_1(X13),X13)|v1_xboole_0(X13))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_24, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, negated_conjecture, (r2_hidden(esk3_0,u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_26, negated_conjecture, (l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, negated_conjecture, (r2_hidden(esk3_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26])]), c_0_27])).
fof(c_0_30, plain, ![X15, X16]:(v1_xboole_0(X15)|~m1_subset_1(X16,X15)|k6_domain_1(X15,X16)=k1_tarski(X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_31, plain, ![X17, X18]:(~r2_hidden(X17,X18)|m1_subset_1(k1_tarski(X17),k1_zfmisc_1(X18))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_subset_1])])).
cnf(c_0_32, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_33, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_35, plain, ![X19, X20]:(~l1_pre_topc(X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|m1_subset_1(k2_pre_topc(X19,X20),k1_zfmisc_1(u1_struct_0(X19)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_pre_topc])])).
cnf(c_0_36, plain, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_37, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_32]), c_0_33])).
fof(c_0_38, plain, ![X21, X22]:(~v2_pre_topc(X21)|~l1_pre_topc(X21)|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))|v4_pre_topc(k2_pre_topc(X21,X22),X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_tops_1])])).
cnf(c_0_39, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk3_0)=k1_tarski(esk3_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_34, c_0_20])).
fof(c_0_40, plain, ![X23, X24, X25, X26, X27, X28]:(((~r2_hidden(X25,X24)|X25=X23|X24!=k1_tarski(X23))&(X26!=X23|r2_hidden(X26,X24)|X24!=k1_tarski(X23)))&((~r2_hidden(esk5_2(X27,X28),X28)|esk5_2(X27,X28)!=X27|X28=k1_tarski(X27))&(r2_hidden(esk5_2(X27,X28),X28)|esk5_2(X27,X28)=X27|X28=k1_tarski(X27)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_41, plain, ![X30, X31, X33, X34, X35]:(((r2_hidden(esk6_2(X30,X31),X30)|r1_xboole_0(X30,X31))&(r2_hidden(esk6_2(X30,X31),X31)|r1_xboole_0(X30,X31)))&(~r2_hidden(X35,X33)|~r2_hidden(X35,X34)|~r1_xboole_0(X33,X34))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_42, plain, ![X36, X37, X38]:(~v2_pre_topc(X36)|~l1_pre_topc(X36)|(~m1_subset_1(X37,k1_zfmisc_1(u1_struct_0(X36)))|(~m1_subset_1(X38,k1_zfmisc_1(u1_struct_0(X36)))|(~r1_xboole_0(X37,X38)|~v3_pre_topc(X37,X36)|r1_xboole_0(X37,k2_pre_topc(X36,X38)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t40_tsep_1])])])).
fof(c_0_43, plain, ![X39, X40]:((~v3_tdlat_3(X39)|(~m1_subset_1(X40,k1_zfmisc_1(u1_struct_0(X39)))|(~v4_pre_topc(X40,X39)|v3_pre_topc(X40,X39)))|(~v2_pre_topc(X39)|~l1_pre_topc(X39)))&((m1_subset_1(esk7_1(X39),k1_zfmisc_1(u1_struct_0(X39)))|v3_tdlat_3(X39)|(~v2_pre_topc(X39)|~l1_pre_topc(X39)))&((v4_pre_topc(esk7_1(X39),X39)|v3_tdlat_3(X39)|(~v2_pre_topc(X39)|~l1_pre_topc(X39)))&(~v3_pre_topc(esk7_1(X39),X39)|v3_tdlat_3(X39)|(~v2_pre_topc(X39)|~l1_pre_topc(X39)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t24_tdlat_3])])])])])).
cnf(c_0_44, plain, (m1_subset_1(k2_pre_topc(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_45, negated_conjecture, (m1_subset_1(k1_tarski(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_46, plain, (v4_pre_topc(k2_pre_topc(X1,X2),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_47, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_48, negated_conjecture, (~r1_xboole_0(k2_pre_topc(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0)),k2_pre_topc(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_49, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk3_0)=k1_tarski(esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_39]), c_0_26])]), c_0_27])).
cnf(c_0_50, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_51, plain, (r2_hidden(esk6_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_52, plain, (r1_xboole_0(X2,k2_pre_topc(X1,X3))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_xboole_0(X2,X3)|~v3_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_53, negated_conjecture, (m1_subset_1(k1_tarski(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_36, c_0_29])).
cnf(c_0_54, plain, (v3_pre_topc(X2,X1)|~v3_tdlat_3(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v4_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_55, negated_conjecture, (m1_subset_1(k2_pre_topc(esk1_0,k1_tarski(esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_22])])).
cnf(c_0_56, negated_conjecture, (v4_pre_topc(k2_pre_topc(esk1_0,k1_tarski(esk2_0)),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_45]), c_0_22]), c_0_47])])).
cnf(c_0_57, negated_conjecture, (v3_tdlat_3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_58, negated_conjecture, (~r1_xboole_0(k2_pre_topc(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0)),k2_pre_topc(esk1_0,k1_tarski(esk3_0)))), inference(rw,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_59, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k1_tarski(esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_32]), c_0_33])).
cnf(c_0_60, plain, (esk6_2(X1,X2)=X3|r1_xboole_0(X1,X2)|X2!=k1_tarski(X3)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
fof(c_0_61, plain, ![X42, X43, X44]:(v2_struct_0(X42)|(~v2_pre_topc(X42)|(~v3_tdlat_3(X42)|~l1_pre_topc(X42)))|(~m1_subset_1(X43,u1_struct_0(X42))|(~m1_subset_1(X44,u1_struct_0(X42))|(~r2_hidden(X43,k2_pre_topc(X42,k6_domain_1(u1_struct_0(X42),X44)))|k2_pre_topc(X42,k6_domain_1(u1_struct_0(X42),X43))=k2_pre_topc(X42,k6_domain_1(u1_struct_0(X42),X44)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t55_tex_2])])])])).
cnf(c_0_62, negated_conjecture, (r1_xboole_0(X1,k2_pre_topc(esk1_0,k1_tarski(esk3_0)))|~v3_pre_topc(X1,esk1_0)|~r1_xboole_0(X1,k1_tarski(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_22]), c_0_47])])).
cnf(c_0_63, negated_conjecture, (v3_pre_topc(k2_pre_topc(esk1_0,k1_tarski(esk2_0)),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56]), c_0_22]), c_0_57]), c_0_47])])).
cnf(c_0_64, negated_conjecture, (~r1_xboole_0(k2_pre_topc(esk1_0,k1_tarski(esk2_0)),k2_pre_topc(esk1_0,k1_tarski(esk3_0)))), inference(rw,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_65, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_66, plain, (esk6_2(X1,k1_tarski(X2))=X2|r1_xboole_0(X1,k1_tarski(X2))), inference(er,[status(thm)],[c_0_60])).
cnf(c_0_67, negated_conjecture, (k2_pre_topc(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=k2_pre_topc(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_68, plain, (v2_struct_0(X1)|k2_pre_topc(X1,k6_domain_1(u1_struct_0(X1),X2))=k2_pre_topc(X1,k6_domain_1(u1_struct_0(X1),X3))|~v2_pre_topc(X1)|~v3_tdlat_3(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X2,k2_pre_topc(X1,k6_domain_1(u1_struct_0(X1),X3)))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_69, negated_conjecture, (~r1_xboole_0(k2_pre_topc(esk1_0,k1_tarski(esk2_0)),k1_tarski(esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_55]), c_0_63])]), c_0_64])).
cnf(c_0_70, plain, (r2_hidden(X1,X2)|r1_xboole_0(X2,k1_tarski(X1))), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_71, negated_conjecture, (k2_pre_topc(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=k2_pre_topc(esk1_0,k1_tarski(esk3_0))), inference(rw,[status(thm)],[c_0_67, c_0_49])).
cnf(c_0_72, negated_conjecture, (k2_pre_topc(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1))=k2_pre_topc(esk1_0,k1_tarski(esk2_0))|~r2_hidden(X1,k2_pre_topc(esk1_0,k1_tarski(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_59]), c_0_32]), c_0_22]), c_0_57]), c_0_47])]), c_0_27])).
cnf(c_0_73, negated_conjecture, (r2_hidden(esk3_0,k2_pre_topc(esk1_0,k1_tarski(esk2_0)))), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_74, negated_conjecture, (k2_pre_topc(esk1_0,k1_tarski(esk3_0))!=k2_pre_topc(esk1_0,k1_tarski(esk2_0))), inference(rw,[status(thm)],[c_0_71, c_0_59])).
cnf(c_0_75, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_49]), c_0_20])]), c_0_74]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 76
# Proof object clause steps            : 47
# Proof object formula steps           : 29
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 38
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 215
# ...of these trivial                  : 1
# ...subsumed                          : 26
# ...remaining for further processing  : 188
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 6
# Generated clauses                    : 445
# ...of the previous two non-trivial   : 389
# Contextual simplify-reflections      : 0
# Paramodulations                      : 438
# Factorizations                       : 0
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 181
#    Positive orientable unit clauses  : 71
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 8
#    Non-unit-clauses                  : 102
# Current number of unprocessed clauses: 194
# ...number of literals in the above   : 564
# Current number of archived formulas  : 0
# Current number of archived clauses   : 6
# Clause-clause subsumption calls (NU) : 3803
# Rec. Clause-clause subsumption calls : 2224
# Non-unit clause-clause subsumptions  : 23
# Unit Clause-clause subsumption calls : 203
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 382
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 11570

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.003 s
# Total time               : 0.017 s
# Maximum resident set size: 28764 pages
