# No SInE strategy applied
# Auto-Mode selected heuristic G_E___200_B02_F1_SE_CS_SP_PI_S2k
# and selection function SelectCQArNTEqFirst.
#
# Preprocessing time       : 0.010 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', t5_subset)).
fof(rc1_funct_2, axiom, ![X1, X2]:?[X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&(v1_relat_1(X3)&(v4_relat_1(X3,X1)&(v5_relat_1(X3,X2)&(v1_funct_1(X3)&v1_funct_2(X3,X1,X2)))))), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', rc1_funct_2)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', d1_xboole_0)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', t113_zfmisc_1)).
fof(t8_boole, axiom, ![X1, X2]:~((v1_xboole_0(X1)&(~(X1=X2)&v1_xboole_0(X2)))), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', t8_boole)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', l13_xboole_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', fc1_xboole_0)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', fc6_relat_1)).
fof(t68_tex_2, conjecture, ![X1]:((v2_pre_topc(X1)&(v1_tdlat_3(X1)&l1_pre_topc(X1)))=>![X2]:((v2_pre_topc(X2)&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>v5_pre_topc(X3,X1,X2)))), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', t68_tex_2)).
fof(d18_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(v4_relat_1(X2,X1)<=>r1_tarski(k1_relat_1(X2),X1))), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', d18_relat_1)).
fof(t17_tdlat_3, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>(v1_tdlat_3(X1)<=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>v3_pre_topc(X2,X1)))), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', t17_tdlat_3)).
fof(cc1_tdlat_3, axiom, ![X1]:(l1_pre_topc(X1)=>(v1_tdlat_3(X1)=>v2_pre_topc(X1))), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', cc1_tdlat_3)).
fof(dt_k8_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k8_relset_1(X1,X2,X3,X4),k1_zfmisc_1(X1))), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', dt_k8_relset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', t3_subset)).
fof(t3_xboole_1, axiom, ![X1]:(r1_tarski(X1,k1_xboole_0)=>X1=k1_xboole_0), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', t3_xboole_1)).
fof(t56_tops_2, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((v2_pre_topc(X2)&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v5_pre_topc(X3,X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>r1_tarski(k2_pre_topc(X1,k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4)),k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,k2_pre_topc(X2,X4)))))))), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', t56_tops_2)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', d4_subset_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', d10_xboole_0)).
fof(t55_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>((k2_struct_0(X2)=k1_xboole_0=>k2_struct_0(X1)=k1_xboole_0)=>(v5_pre_topc(X3,X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(v3_pre_topc(X4,X2)=>v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X1)))))))), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', t55_tops_2)).
fof(rc1_partfun1, axiom, ![X1, X2]:?[X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&(v1_relat_1(X3)&(v4_relat_1(X3,X1)&(v5_relat_1(X3,X2)&v1_funct_1(X3))))), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', rc1_partfun1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', redefinition_k1_relset_1)).
fof(cc4_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', cc4_relset_1)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', d3_struct_0)).
fof(t38_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(k7_relset_1(X1,X2,X3,X1)=k2_relset_1(X1,X2,X3)&k8_relset_1(X1,X2,X3,X2)=k1_relset_1(X1,X2,X3))), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', t38_relset_1)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', dt_l1_pre_topc)).
fof(t18_tdlat_3, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>(v1_tdlat_3(X1)<=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>v4_pre_topc(X2,X1)))), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', t18_tdlat_3)).
fof(t52_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v4_pre_topc(X2,X1)=>k2_pre_topc(X1,X2)=X2)&((v2_pre_topc(X1)&k2_pre_topc(X1,X2)=X2)=>v4_pre_topc(X2,X1))))), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', t52_pre_topc)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/76796084cd774b3f9fc17c475d822838.E_input', t2_xboole_1)).
fof(c_0_29, plain, ![X5, X6, X7]:(~r2_hidden(X5,X6)|(~m1_subset_1(X6,k1_zfmisc_1(X7))|~v1_xboole_0(X7))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_30, plain, ![X8, X9]:(m1_subset_1(esk1_2(X8,X9),k1_zfmisc_1(k2_zfmisc_1(X8,X9)))&(v1_relat_1(esk1_2(X8,X9))&(v4_relat_1(esk1_2(X8,X9),X8)&(v5_relat_1(esk1_2(X8,X9),X9)&(v1_funct_1(esk1_2(X8,X9))&v1_funct_2(esk1_2(X8,X9),X8,X9)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc1_funct_2])])).
cnf(c_0_31, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_32, plain, (m1_subset_1(esk1_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_33, plain, ![X11, X12, X13]:((~v1_xboole_0(X11)|~r2_hidden(X12,X11))&(r2_hidden(esk2_1(X13),X13)|v1_xboole_0(X13))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_34, plain, ![X15, X16]:((k2_zfmisc_1(X15,X16)!=k1_xboole_0|(X15=k1_xboole_0|X16=k1_xboole_0))&((X15!=k1_xboole_0|k2_zfmisc_1(X15,X16)=k1_xboole_0)&(X16!=k1_xboole_0|k2_zfmisc_1(X15,X16)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_35, plain, (~v1_xboole_0(k2_zfmisc_1(X1,X2))|~r2_hidden(X3,esk1_2(X1,X2))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_36, plain, (r2_hidden(esk2_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_37, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_38, plain, ![X17, X18]:(~v1_xboole_0(X17)|(X17=X18|~v1_xboole_0(X18))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_boole])])])).
fof(c_0_39, plain, ![X19]:(~v1_xboole_0(X19)|X19=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_40, plain, (v1_xboole_0(esk1_2(X1,X2))|~v1_xboole_0(k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_41, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_42, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
fof(c_0_43, plain, ![X20, X21]:v1_relat_1(k2_zfmisc_1(X20,X21)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_44, plain, (v4_relat_1(esk1_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_45, plain, (X1=X2|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_46, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_47, plain, (v1_xboole_0(esk1_2(X1,k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42])])).
cnf(c_0_48, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_50, negated_conjecture, ~(![X1]:((v2_pre_topc(X1)&(v1_tdlat_3(X1)&l1_pre_topc(X1)))=>![X2]:((v2_pre_topc(X2)&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>v5_pre_topc(X3,X1,X2))))), inference(assume_negation,[status(cth)],[t68_tex_2])).
fof(c_0_51, plain, ![X25, X26]:((~v4_relat_1(X26,X25)|r1_tarski(k1_relat_1(X26),X25)|~v1_relat_1(X26))&(~r1_tarski(k1_relat_1(X26),X25)|v4_relat_1(X26,X25)|~v1_relat_1(X26))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d18_relat_1])])])).
cnf(c_0_52, plain, (v4_relat_1(X1,X2)|~v1_xboole_0(esk1_2(X2,X3))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_53, plain, (esk1_2(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_54, plain, (v1_relat_1(X1)|~v1_xboole_0(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_48, c_0_45])).
cnf(c_0_55, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_49])).
fof(c_0_56, plain, ![X27, X28]:((~v1_tdlat_3(X27)|(~m1_subset_1(X28,k1_zfmisc_1(u1_struct_0(X27)))|v3_pre_topc(X28,X27))|(~v2_pre_topc(X27)|~l1_pre_topc(X27)))&((m1_subset_1(esk6_1(X27),k1_zfmisc_1(u1_struct_0(X27)))|v1_tdlat_3(X27)|(~v2_pre_topc(X27)|~l1_pre_topc(X27)))&(~v3_pre_topc(esk6_1(X27),X27)|v1_tdlat_3(X27)|(~v2_pre_topc(X27)|~l1_pre_topc(X27))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t17_tdlat_3])])])])])).
fof(c_0_57, plain, ![X30]:(~l1_pre_topc(X30)|(~v1_tdlat_3(X30)|v2_pre_topc(X30))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_tdlat_3])])).
fof(c_0_58, plain, ![X31, X32, X33, X34]:(~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X32)))|m1_subset_1(k8_relset_1(X31,X32,X33,X34),k1_zfmisc_1(X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_relset_1])])).
fof(c_0_59, negated_conjecture, ((v2_pre_topc(esk3_0)&(v1_tdlat_3(esk3_0)&l1_pre_topc(esk3_0)))&((v2_pre_topc(esk4_0)&l1_pre_topc(esk4_0))&((v1_funct_1(esk5_0)&(v1_funct_2(esk5_0,u1_struct_0(esk3_0),u1_struct_0(esk4_0))&m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk4_0))))))&~v5_pre_topc(esk5_0,esk3_0,esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_50])])])).
fof(c_0_60, plain, ![X35, X36]:((~m1_subset_1(X35,k1_zfmisc_1(X36))|r1_tarski(X35,X36))&(~r1_tarski(X35,X36)|m1_subset_1(X35,k1_zfmisc_1(X36)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_61, plain, (r1_tarski(k1_relat_1(X1),X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_62, plain, (v4_relat_1(X1,X2)|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_42])])).
cnf(c_0_63, plain, (v1_relat_1(X1)|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_42])])).
fof(c_0_64, plain, ![X37]:(~r1_tarski(X37,k1_xboole_0)|X37=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_xboole_1])])).
cnf(c_0_65, plain, (v4_relat_1(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_44, c_0_53])).
cnf(c_0_66, plain, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_48, c_0_41])).
cnf(c_0_67, plain, (v3_pre_topc(X2,X1)|~v1_tdlat_3(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_68, plain, (v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_69, plain, (m1_subset_1(k8_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_70, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk4_0))))), inference(split_conjunct,[status(thm)],[c_0_59])).
fof(c_0_71, plain, ![X56, X57, X58, X59]:((~v5_pre_topc(X58,X56,X57)|(~m1_subset_1(X59,k1_zfmisc_1(u1_struct_0(X57)))|r1_tarski(k2_pre_topc(X56,k8_relset_1(u1_struct_0(X56),u1_struct_0(X57),X58,X59)),k8_relset_1(u1_struct_0(X56),u1_struct_0(X57),X58,k2_pre_topc(X57,X59))))|(~v1_funct_1(X58)|(~v1_funct_2(X58,u1_struct_0(X56),u1_struct_0(X57))|~m1_subset_1(X58,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X56),u1_struct_0(X57))))))|(~v2_pre_topc(X57)|~l1_pre_topc(X57))|(~v2_pre_topc(X56)|~l1_pre_topc(X56)))&((m1_subset_1(esk9_3(X56,X57,X58),k1_zfmisc_1(u1_struct_0(X57)))|v5_pre_topc(X58,X56,X57)|(~v1_funct_1(X58)|(~v1_funct_2(X58,u1_struct_0(X56),u1_struct_0(X57))|~m1_subset_1(X58,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X56),u1_struct_0(X57))))))|(~v2_pre_topc(X57)|~l1_pre_topc(X57))|(~v2_pre_topc(X56)|~l1_pre_topc(X56)))&(~r1_tarski(k2_pre_topc(X56,k8_relset_1(u1_struct_0(X56),u1_struct_0(X57),X58,esk9_3(X56,X57,X58))),k8_relset_1(u1_struct_0(X56),u1_struct_0(X57),X58,k2_pre_topc(X57,esk9_3(X56,X57,X58))))|v5_pre_topc(X58,X56,X57)|(~v1_funct_1(X58)|(~v1_funct_2(X58,u1_struct_0(X56),u1_struct_0(X57))|~m1_subset_1(X58,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X56),u1_struct_0(X57))))))|(~v2_pre_topc(X57)|~l1_pre_topc(X57))|(~v2_pre_topc(X56)|~l1_pre_topc(X56))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t56_tops_2])])])])])).
fof(c_0_72, plain, ![X38]:m1_subset_1(k2_subset_1(X38),k1_zfmisc_1(X38)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_73, plain, ![X39]:k2_subset_1(X39)=X39, inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_74, plain, ![X43, X44]:(((r1_tarski(X43,X44)|X43!=X44)&(r1_tarski(X44,X43)|X43!=X44))&(~r1_tarski(X43,X44)|~r1_tarski(X44,X43)|X43=X44)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_75, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_76, plain, (r1_tarski(k1_relat_1(X1),X2)|~v1_xboole_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63])).
cnf(c_0_77, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_78, plain, (r1_tarski(k1_relat_1(k1_xboole_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_65]), c_0_66])])).
fof(c_0_79, plain, ![X45, X46, X47, X48]:(((~v5_pre_topc(X47,X45,X46)|(~m1_subset_1(X48,k1_zfmisc_1(u1_struct_0(X46)))|(~v3_pre_topc(X48,X46)|v3_pre_topc(k8_relset_1(u1_struct_0(X45),u1_struct_0(X46),X47,X48),X45)))|k2_struct_0(X46)=k1_xboole_0|(~v1_funct_1(X47)|(~v1_funct_2(X47,u1_struct_0(X45),u1_struct_0(X46))|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X45),u1_struct_0(X46))))))|~l1_pre_topc(X46)|~l1_pre_topc(X45))&((m1_subset_1(esk8_3(X45,X46,X47),k1_zfmisc_1(u1_struct_0(X46)))|v5_pre_topc(X47,X45,X46)|k2_struct_0(X46)=k1_xboole_0|(~v1_funct_1(X47)|(~v1_funct_2(X47,u1_struct_0(X45),u1_struct_0(X46))|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X45),u1_struct_0(X46))))))|~l1_pre_topc(X46)|~l1_pre_topc(X45))&((v3_pre_topc(esk8_3(X45,X46,X47),X46)|v5_pre_topc(X47,X45,X46)|k2_struct_0(X46)=k1_xboole_0|(~v1_funct_1(X47)|(~v1_funct_2(X47,u1_struct_0(X45),u1_struct_0(X46))|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X45),u1_struct_0(X46))))))|~l1_pre_topc(X46)|~l1_pre_topc(X45))&(~v3_pre_topc(k8_relset_1(u1_struct_0(X45),u1_struct_0(X46),X47,esk8_3(X45,X46,X47)),X45)|v5_pre_topc(X47,X45,X46)|k2_struct_0(X46)=k1_xboole_0|(~v1_funct_1(X47)|(~v1_funct_2(X47,u1_struct_0(X45),u1_struct_0(X46))|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X45),u1_struct_0(X46))))))|~l1_pre_topc(X46)|~l1_pre_topc(X45)))))&((~v5_pre_topc(X47,X45,X46)|(~m1_subset_1(X48,k1_zfmisc_1(u1_struct_0(X46)))|(~v3_pre_topc(X48,X46)|v3_pre_topc(k8_relset_1(u1_struct_0(X45),u1_struct_0(X46),X47,X48),X45)))|k2_struct_0(X45)!=k1_xboole_0|(~v1_funct_1(X47)|(~v1_funct_2(X47,u1_struct_0(X45),u1_struct_0(X46))|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X45),u1_struct_0(X46))))))|~l1_pre_topc(X46)|~l1_pre_topc(X45))&((m1_subset_1(esk8_3(X45,X46,X47),k1_zfmisc_1(u1_struct_0(X46)))|v5_pre_topc(X47,X45,X46)|k2_struct_0(X45)!=k1_xboole_0|(~v1_funct_1(X47)|(~v1_funct_2(X47,u1_struct_0(X45),u1_struct_0(X46))|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X45),u1_struct_0(X46))))))|~l1_pre_topc(X46)|~l1_pre_topc(X45))&((v3_pre_topc(esk8_3(X45,X46,X47),X46)|v5_pre_topc(X47,X45,X46)|k2_struct_0(X45)!=k1_xboole_0|(~v1_funct_1(X47)|(~v1_funct_2(X47,u1_struct_0(X45),u1_struct_0(X46))|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X45),u1_struct_0(X46))))))|~l1_pre_topc(X46)|~l1_pre_topc(X45))&(~v3_pre_topc(k8_relset_1(u1_struct_0(X45),u1_struct_0(X46),X47,esk8_3(X45,X46,X47)),X45)|v5_pre_topc(X47,X45,X46)|k2_struct_0(X45)!=k1_xboole_0|(~v1_funct_1(X47)|(~v1_funct_2(X47,u1_struct_0(X45),u1_struct_0(X46))|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X45),u1_struct_0(X46))))))|~l1_pre_topc(X46)|~l1_pre_topc(X45)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_tops_2])])])])])).
cnf(c_0_80, plain, (v3_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v1_tdlat_3(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))), inference(csr,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_81, negated_conjecture, (m1_subset_1(k8_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk4_0),esk5_0,X1),k1_zfmisc_1(u1_struct_0(esk3_0)))), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_82, negated_conjecture, (l1_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_83, negated_conjecture, (v1_tdlat_3(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_84, plain, (m1_subset_1(esk9_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|v5_pre_topc(X3,X1,X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_85, negated_conjecture, (v1_funct_2(esk5_0,u1_struct_0(esk3_0),u1_struct_0(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_86, negated_conjecture, (l1_pre_topc(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_87, negated_conjecture, (v2_pre_topc(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_88, negated_conjecture, (v2_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_89, negated_conjecture, (v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_90, negated_conjecture, (~v5_pre_topc(esk5_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
fof(c_0_91, plain, ![X40, X41]:(m1_subset_1(esk7_2(X40,X41),k1_zfmisc_1(k2_zfmisc_1(X40,X41)))&(v1_relat_1(esk7_2(X40,X41))&(v4_relat_1(esk7_2(X40,X41),X40)&(v5_relat_1(esk7_2(X40,X41),X41)&v1_funct_1(esk7_2(X40,X41)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc1_partfun1])])).
fof(c_0_92, plain, ![X50, X51, X52]:(~m1_subset_1(X52,k1_zfmisc_1(k2_zfmisc_1(X50,X51)))|k1_relset_1(X50,X51,X52)=k1_relat_1(X52)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
fof(c_0_93, plain, ![X53, X54, X55]:(~v1_xboole_0(X53)|(~m1_subset_1(X55,k1_zfmisc_1(k2_zfmisc_1(X54,X53)))|v1_xboole_0(X55))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])).
cnf(c_0_94, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_95, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_96, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_97, negated_conjecture, (r1_tarski(esk5_0,k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk4_0)))), inference(spm,[status(thm)],[c_0_75, c_0_70])).
cnf(c_0_98, plain, (r1_tarski(X1,X2)|~v1_xboole_0(k1_relat_1(X3))|~v1_xboole_0(X3)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_76, c_0_45])).
cnf(c_0_99, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_77, c_0_78])).
fof(c_0_100, plain, ![X61]:(~l1_struct_0(X61)|k2_struct_0(X61)=u1_struct_0(X61)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_101, plain, (v5_pre_topc(X3,X1,X2)|k2_struct_0(X2)=k1_xboole_0|~v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk8_3(X1,X2,X3)),X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_102, negated_conjecture, (v3_pre_topc(k8_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk4_0),esk5_0,X1),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_82]), c_0_83])])).
cnf(c_0_103, negated_conjecture, (m1_subset_1(esk9_3(esk3_0,esk4_0,esk5_0),k1_zfmisc_1(u1_struct_0(esk4_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86]), c_0_82]), c_0_87]), c_0_88]), c_0_89]), c_0_70])]), c_0_90])).
cnf(c_0_104, plain, (m1_subset_1(esk7_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_91])).
fof(c_0_105, plain, ![X62, X63, X64]:((k7_relset_1(X62,X63,X64,X62)=k2_relset_1(X62,X63,X64)|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(X62,X63))))&(k8_relset_1(X62,X63,X64,X63)=k1_relset_1(X62,X63,X64)|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(X62,X63))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t38_relset_1])])])).
cnf(c_0_106, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_107, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_108, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_109, negated_conjecture, (k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk4_0))=esk5_0|~r1_tarski(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk4_0)),esk5_0)), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_110, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_42])])).
cnf(c_0_111, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_112, negated_conjecture, (k2_struct_0(esk4_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_85]), c_0_86]), c_0_82]), c_0_89]), c_0_70])]), c_0_90]), c_0_102])])).
fof(c_0_113, plain, ![X65]:(~l1_pre_topc(X65)|l1_struct_0(X65)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_114, negated_conjecture, (r1_tarski(esk9_3(esk3_0,esk4_0,esk5_0),u1_struct_0(esk4_0))), inference(spm,[status(thm)],[c_0_75, c_0_103])).
cnf(c_0_115, plain, (~v1_xboole_0(k2_zfmisc_1(X1,X2))|~r2_hidden(X3,esk7_2(X1,X2))), inference(spm,[status(thm)],[c_0_31, c_0_104])).
fof(c_0_116, plain, ![X66, X67]:((~v1_tdlat_3(X66)|(~m1_subset_1(X67,k1_zfmisc_1(u1_struct_0(X66)))|v4_pre_topc(X67,X66))|(~v2_pre_topc(X66)|~l1_pre_topc(X66)))&((m1_subset_1(esk10_1(X66),k1_zfmisc_1(u1_struct_0(X66)))|v1_tdlat_3(X66)|(~v2_pre_topc(X66)|~l1_pre_topc(X66)))&(~v4_pre_topc(esk10_1(X66),X66)|v1_tdlat_3(X66)|(~v2_pre_topc(X66)|~l1_pre_topc(X66))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t18_tdlat_3])])])])])).
cnf(c_0_117, plain, (k8_relset_1(X1,X2,X3,X2)=k1_relset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_118, negated_conjecture, (k1_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk4_0),esk5_0)=k1_relat_1(esk5_0)), inference(spm,[status(thm)],[c_0_106, c_0_70])).
cnf(c_0_119, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_120, negated_conjecture, (k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk4_0))=esk5_0|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk4_0)))), inference(spm,[status(thm)],[c_0_109, c_0_110])).
cnf(c_0_121, negated_conjecture, (u1_struct_0(esk4_0)=k1_xboole_0|~l1_struct_0(esk4_0)), inference(spm,[status(thm)],[c_0_111, c_0_112])).
cnf(c_0_122, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_123, negated_conjecture, (esk9_3(esk3_0,esk4_0,esk5_0)=u1_struct_0(esk4_0)|~r1_tarski(u1_struct_0(esk4_0),esk9_3(esk3_0,esk4_0,esk5_0))), inference(spm,[status(thm)],[c_0_96, c_0_114])).
cnf(c_0_124, plain, (v1_xboole_0(esk7_2(X1,X2))|~v1_xboole_0(k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_115, c_0_36])).
cnf(c_0_125, plain, (v4_pre_topc(X2,X1)|~v1_tdlat_3(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_116])).
cnf(c_0_126, negated_conjecture, (k8_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk4_0),esk5_0,u1_struct_0(esk4_0))=k1_relat_1(esk5_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_70]), c_0_118])).
cnf(c_0_127, plain, (v1_xboole_0(esk1_2(X1,X2))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_40, c_0_119])).
cnf(c_0_128, plain, (v1_funct_1(esk1_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_129, plain, (v1_xboole_0(esk1_2(k1_xboole_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_55]), c_0_42])])).
cnf(c_0_130, negated_conjecture, (k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk4_0))=esk5_0|~v1_xboole_0(u1_struct_0(esk4_0))), inference(spm,[status(thm)],[c_0_120, c_0_119])).
cnf(c_0_131, negated_conjecture, (u1_struct_0(esk4_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_86])])).
cnf(c_0_132, negated_conjecture, (esk9_3(esk3_0,esk4_0,esk5_0)=u1_struct_0(esk4_0)|~v1_xboole_0(u1_struct_0(esk4_0))), inference(spm,[status(thm)],[c_0_123, c_0_110])).
cnf(c_0_133, negated_conjecture, (~v1_xboole_0(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk4_0)))|~r2_hidden(X1,esk5_0)), inference(spm,[status(thm)],[c_0_31, c_0_70])).
cnf(c_0_134, plain, (v1_xboole_0(esk7_2(X1,k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_41]), c_0_42])])).
fof(c_0_135, plain, ![X69, X70]:((~v4_pre_topc(X70,X69)|k2_pre_topc(X69,X70)=X70|~m1_subset_1(X70,k1_zfmisc_1(u1_struct_0(X69)))|~l1_pre_topc(X69))&(~v2_pre_topc(X69)|k2_pre_topc(X69,X70)!=X70|v4_pre_topc(X70,X69)|~m1_subset_1(X70,k1_zfmisc_1(u1_struct_0(X69)))|~l1_pre_topc(X69))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t52_pre_topc])])])])).
cnf(c_0_136, plain, (v4_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v1_tdlat_3(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))), inference(csr,[status(thm)],[c_0_125, c_0_68])).
cnf(c_0_137, negated_conjecture, (m1_subset_1(k1_relat_1(esk5_0),k1_zfmisc_1(u1_struct_0(esk3_0)))), inference(spm,[status(thm)],[c_0_81, c_0_126])).
cnf(c_0_138, plain, (v1_funct_2(esk1_2(X1,X2),X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_139, plain, (esk1_2(X1,X2)=X3|~v1_xboole_0(X3)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_45, c_0_127])).
cnf(c_0_140, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_141, plain, (v1_funct_1(X1)|~v1_xboole_0(esk1_2(X2,X3))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_128, c_0_45])).
cnf(c_0_142, plain, (esk1_2(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_46, c_0_129])).
cnf(c_0_143, negated_conjecture, (k1_xboole_0=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_130, c_0_131]), c_0_41]), c_0_131]), c_0_42])])).
cnf(c_0_144, negated_conjecture, (esk9_3(esk3_0,esk4_0,esk5_0)=X1|~v1_xboole_0(u1_struct_0(esk4_0))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_132, c_0_45])).
cnf(c_0_145, negated_conjecture, (v1_xboole_0(esk5_0)|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk4_0)))), inference(spm,[status(thm)],[c_0_133, c_0_36])).
cnf(c_0_146, plain, (k1_relset_1(X1,X2,esk7_2(X1,X2))=k1_relat_1(esk7_2(X1,X2))), inference(spm,[status(thm)],[c_0_106, c_0_104])).
cnf(c_0_147, plain, (esk7_2(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_46, c_0_134])).
cnf(c_0_148, plain, (k2_pre_topc(X2,X1)=X1|~v4_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_135])).
cnf(c_0_149, negated_conjecture, (v4_pre_topc(k1_relat_1(esk5_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_137]), c_0_82]), c_0_83])])).
fof(c_0_150, plain, ![X71]:r1_tarski(k1_xboole_0,X71), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_151, plain, (v5_pre_topc(X3,X1,X2)|~r1_tarski(k2_pre_topc(X1,k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk9_3(X1,X2,X3))),k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,k2_pre_topc(X2,esk9_3(X1,X2,X3))))|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_152, plain, (v1_funct_2(X1,X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X3)), inference(spm,[status(thm)],[c_0_138, c_0_139])).
cnf(c_0_153, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_140, c_0_110])).
cnf(c_0_154, plain, (v1_funct_1(X1)|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_142]), c_0_42])])).
cnf(c_0_155, plain, (k2_zfmisc_1(esk5_0,X1)=esk5_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_143]), c_0_143])).
cnf(c_0_156, negated_conjecture, (esk9_3(esk3_0,esk4_0,esk5_0)=X1|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_144, c_0_131]), c_0_42])])).
cnf(c_0_157, negated_conjecture, (v1_xboole_0(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_145, c_0_131]), c_0_41]), c_0_42])])).
cnf(c_0_158, plain, (k8_relset_1(X1,X2,esk7_2(X1,X2),X2)=k1_relat_1(esk7_2(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_104]), c_0_146])).
cnf(c_0_159, plain, (esk7_2(X1,esk5_0)=esk5_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_147, c_0_143]), c_0_143])).
cnf(c_0_160, plain, (k1_relat_1(esk5_0)=esk5_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_99, c_0_143]), c_0_143])).
cnf(c_0_161, negated_conjecture, (k2_pre_topc(esk3_0,k1_relat_1(esk5_0))=k1_relat_1(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_137]), c_0_82])]), c_0_149])])).
cnf(c_0_162, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_150])).
cnf(c_0_163, plain, (v5_pre_topc(X1,X2,X3)|~r1_tarski(k2_pre_topc(X2,k8_relset_1(u1_struct_0(X2),u1_struct_0(X3),X1,esk9_3(X2,X3,X1))),k8_relset_1(u1_struct_0(X2),u1_struct_0(X3),X1,k2_pre_topc(X3,esk9_3(X2,X3,X1))))|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~v2_pre_topc(X2)|~v1_xboole_0(u1_struct_0(X3))|~v1_xboole_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152]), c_0_153]), c_0_154])).
cnf(c_0_164, negated_conjecture, (esk9_3(esk3_0,esk4_0,esk5_0)=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_156]), c_0_155]), c_0_157])])).
cnf(c_0_165, plain, (k8_relset_1(X1,esk5_0,esk5_0,esk5_0)=esk5_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159]), c_0_160])).
cnf(c_0_166, negated_conjecture, (k2_pre_topc(esk3_0,esk5_0)=esk5_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_161, c_0_160]), c_0_160])).
cnf(c_0_167, plain, (r1_tarski(esk5_0,X1)), inference(rw,[status(thm)],[c_0_162, c_0_143])).
cnf(c_0_168, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_164]), c_0_131]), c_0_143]), c_0_165]), c_0_166]), c_0_131]), c_0_143]), c_0_167]), c_0_86]), c_0_82]), c_0_87]), c_0_88]), c_0_131]), c_0_143]), c_0_157]), c_0_157])]), c_0_90]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 169
# Proof object clause steps            : 111
# Proof object formula steps           : 58
# Proof object conjectures             : 39
# Proof object clause conjectures      : 36
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 43
# Proof object initial formulas used   : 29
# Proof object generating inferences   : 55
# Proof object simplifying inferences  : 96
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 29
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 68
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 67
# Processed clauses                    : 2763
# ...of these trivial                  : 79
# ...subsumed                          : 1963
# ...remaining for further processing  : 721
# Other redundant clauses eliminated   : 32
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 59
# Backward-rewritten                   : 167
# Generated clauses                    : 9319
# ...of the previous two non-trivial   : 8591
# Contextual simplify-reflections      : 23
# Paramodulations                      : 9287
# Factorizations                       : 0
# Equation resolutions                 : 32
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 491
#    Positive orientable unit clauses  : 72
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 417
# Current number of unprocessed clauses: 5412
# ...number of literals in the above   : 28757
# Current number of archived formulas  : 0
# Current number of archived clauses   : 227
# Clause-clause subsumption calls (NU) : 114165
# Rec. Clause-clause subsumption calls : 52410
# Non-unit clause-clause subsumptions  : 1875
# Unit Clause-clause subsumption calls : 4073
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 133
# BW rewrite match successes           : 24
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 131461

# -------------------------------------------------
# User time                : 0.115 s
# System time              : 0.007 s
# Total time               : 0.122 s
# Maximum resident set size: 28756 pages
