# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t17_tdlat_3, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>(v1_tdlat_3(X1)<=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>v3_pre_topc(X2,X1)))), file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', t17_tdlat_3)).
fof(cc1_tdlat_3, axiom, ![X1]:(l1_pre_topc(X1)=>(v1_tdlat_3(X1)=>v2_pre_topc(X1))), file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', cc1_tdlat_3)).
fof(t49_tex_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v1_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_tex_2(X2,X1)<=>~(v1_subset_1(X2,u1_struct_0(X1)))))), file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', t49_tex_2)).
fof(d7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(v1_subset_1(X2,X1)<=>~(X2=X1))), file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', d7_subset_1)).
fof(t47_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v3_pre_topc(X2,X1)&v3_tex_2(X2,X1))=>v1_tops_1(X2,X1)))), file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', t47_tex_2)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', dt_l1_pre_topc)).
fof(fc12_tops_1, axiom, ![X1]:(l1_pre_topc(X1)=>v1_tops_1(k2_struct_0(X1),X1)), file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', fc12_tops_1)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', d3_struct_0)).
fof(d2_tops_3, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v1_tops_1(X2,X1)<=>k2_pre_topc(X1,X2)=u1_struct_0(X1)))), file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', d2_tops_3)).
fof(t43_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v1_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>v2_tex_2(X2,X1))), file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', t43_tex_2)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', t2_subset)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', fc1_subset_1)).
fof(dt_k3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>m1_subset_1(k3_subset_1(X1,X2),k1_zfmisc_1(X1))), file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', dt_k3_subset_1)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', d1_zfmisc_1)).
fof(t29_tops_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v4_pre_topc(X2,X1)<=>v3_pre_topc(k3_subset_1(u1_struct_0(X1),X2),X1)))), file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', t29_tops_1)).
fof(t48_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v1_tops_1(X2,X1)&v2_tex_2(X2,X1))=>v3_tex_2(X2,X1)))), file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', t48_tex_2)).
fof(t31_tops_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((v4_pre_topc(X2,X1)&r1_tarski(X3,X2))=>r1_tarski(k2_pre_topc(X1,X3),X2))))), file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', t31_tops_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', d10_xboole_0)).
fof(fc14_subset_1, axiom, ![X1]:~(v1_subset_1(k2_subset_1(X1),X1)), file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', fc14_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/76a81de8bfca44a79da3baf1ab54c645.E_input', d4_subset_1)).
fof(c_0_20, plain, ![X24, X25]:((~v1_tdlat_3(X24)|(~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|v3_pre_topc(X25,X24))|(~v2_pre_topc(X24)|~l1_pre_topc(X24)))&((m1_subset_1(esk3_1(X24),k1_zfmisc_1(u1_struct_0(X24)))|v1_tdlat_3(X24)|(~v2_pre_topc(X24)|~l1_pre_topc(X24)))&(~v3_pre_topc(esk3_1(X24),X24)|v1_tdlat_3(X24)|(~v2_pre_topc(X24)|~l1_pre_topc(X24))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t17_tdlat_3])])])])])).
fof(c_0_21, plain, ![X6]:(~l1_pre_topc(X6)|(~v1_tdlat_3(X6)|v2_pre_topc(X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_tdlat_3])])).
fof(c_0_22, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v1_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_tex_2(X2,X1)<=>~(v1_subset_1(X2,u1_struct_0(X1))))))), inference(assume_negation,[status(cth)],[t49_tex_2])).
cnf(c_0_23, plain, (v3_pre_topc(X2,X1)|~v1_tdlat_3(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24, plain, (v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_25, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&(v1_tdlat_3(esk1_0)&l1_pre_topc(esk1_0))))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&((~v3_tex_2(esk2_0,esk1_0)|v1_subset_1(esk2_0,u1_struct_0(esk1_0)))&(v3_tex_2(esk2_0,esk1_0)|~v1_subset_1(esk2_0,u1_struct_0(esk1_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_22])])])])).
fof(c_0_26, plain, ![X35, X36]:((~v1_subset_1(X36,X35)|X36!=X35|~m1_subset_1(X36,k1_zfmisc_1(X35)))&(X36=X35|v1_subset_1(X36,X35)|~m1_subset_1(X36,k1_zfmisc_1(X35)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_subset_1])])])])).
fof(c_0_27, plain, ![X7, X8]:(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|(~v3_pre_topc(X8,X7)|~v3_tex_2(X8,X7)|v1_tops_1(X8,X7)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t47_tex_2])])])])).
cnf(c_0_28, plain, (v3_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v1_tdlat_3(X2)), inference(csr,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, negated_conjecture, (v1_tdlat_3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, plain, (X1=X2|v1_subset_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, plain, (v2_struct_0(X1)|v1_tops_1(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_pre_topc(X2,X1)|~v3_tex_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_34, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_36, negated_conjecture, (v3_pre_topc(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30]), c_0_31])])).
cnf(c_0_37, negated_conjecture, (v3_tex_2(esk2_0,esk1_0)|~v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_38, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_32, c_0_29])).
fof(c_0_39, plain, ![X27]:(~l1_pre_topc(X27)|l1_struct_0(X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_40, plain, ![X21]:(~l1_pre_topc(X21)|v1_tops_1(k2_struct_0(X21),X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc12_tops_1])])).
fof(c_0_41, plain, ![X22]:(~l1_struct_0(X22)|k2_struct_0(X22)=u1_struct_0(X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_42, negated_conjecture, (v1_tops_1(esk2_0,esk1_0)|~v3_tex_2(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_29]), c_0_30]), c_0_34])]), c_0_35]), c_0_36])])).
cnf(c_0_43, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|v3_tex_2(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_44, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_45, plain, ![X18, X19]:((~v1_tops_1(X19,X18)|k2_pre_topc(X18,X19)=u1_struct_0(X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|~l1_pre_topc(X18))&(k2_pre_topc(X18,X19)!=u1_struct_0(X18)|v1_tops_1(X19,X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|~l1_pre_topc(X18))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tops_3])])])])).
cnf(c_0_46, plain, (v1_tops_1(k2_struct_0(X1),X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_47, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_48, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|v1_tops_1(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_49, negated_conjecture, (l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_44, c_0_30])).
cnf(c_0_50, plain, (k2_pre_topc(X2,X1)=u1_struct_0(X2)|~v1_tops_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_51, negated_conjecture, (v1_tops_1(k2_struct_0(esk1_0),esk1_0)), inference(spm,[status(thm)],[c_0_46, c_0_30])).
cnf(c_0_52, negated_conjecture, (k2_struct_0(esk1_0)=esk2_0|v1_tops_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])])).
cnf(c_0_53, negated_conjecture, (u1_struct_0(esk1_0)=k2_pre_topc(esk1_0,esk2_0)|~v1_tops_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_29]), c_0_30])])).
cnf(c_0_54, negated_conjecture, (v1_tops_1(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
fof(c_0_55, plain, ![X11, X12]:(v2_struct_0(X11)|(~v2_pre_topc(X11)|(~v1_tdlat_3(X11)|~l1_pre_topc(X11)))|(~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|v2_tex_2(X12,X11))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t43_tex_2])])])])).
fof(c_0_56, plain, ![X33, X34]:(~m1_subset_1(X33,X34)|(v1_xboole_0(X34)|r2_hidden(X33,X34))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_57, plain, ![X23]:~v1_xboole_0(k1_zfmisc_1(X23)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
cnf(c_0_58, negated_conjecture, (u1_struct_0(esk1_0)=k2_pre_topc(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_54])])).
fof(c_0_59, plain, ![X31, X32]:(~m1_subset_1(X32,k1_zfmisc_1(X31))|m1_subset_1(k3_subset_1(X31,X32),k1_zfmisc_1(X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_subset_1])])).
cnf(c_0_60, plain, (v2_struct_0(X1)|v2_tex_2(X2,X1)|~v2_pre_topc(X1)|~v1_tdlat_3(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_55])).
fof(c_0_61, plain, ![X37, X38, X39, X40, X41, X42]:(((~r2_hidden(X39,X38)|r1_tarski(X39,X37)|X38!=k1_zfmisc_1(X37))&(~r1_tarski(X40,X37)|r2_hidden(X40,X38)|X38!=k1_zfmisc_1(X37)))&((~r2_hidden(esk4_2(X41,X42),X42)|~r1_tarski(esk4_2(X41,X42),X41)|X42=k1_zfmisc_1(X41))&(r2_hidden(esk4_2(X41,X42),X42)|r1_tarski(esk4_2(X41,X42),X41)|X42=k1_zfmisc_1(X41)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
cnf(c_0_62, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_63, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_57])).
fof(c_0_64, plain, ![X9, X10]:((~v4_pre_topc(X10,X9)|v3_pre_topc(k3_subset_1(u1_struct_0(X9),X10),X9)|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|~l1_pre_topc(X9))&(~v3_pre_topc(k3_subset_1(u1_struct_0(X9),X10),X9)|v4_pre_topc(X10,X9)|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|~l1_pre_topc(X9))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t29_tops_1])])])])).
cnf(c_0_65, negated_conjecture, (v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_pre_topc(esk1_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_58]), c_0_30]), c_0_31])])).
cnf(c_0_66, plain, (m1_subset_1(k3_subset_1(X2,X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_59])).
fof(c_0_67, plain, ![X28, X29]:(v2_struct_0(X28)|(~v2_pre_topc(X28)|~l1_pre_topc(X28))|(~m1_subset_1(X29,k1_zfmisc_1(u1_struct_0(X28)))|(~v1_tops_1(X29,X28)|~v2_tex_2(X29,X28)|v3_tex_2(X29,X28)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t48_tex_2])])])])).
cnf(c_0_68, plain, (v2_tex_2(X1,X2)|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v1_tdlat_3(X2)), inference(csr,[status(thm)],[c_0_60, c_0_24])).
cnf(c_0_69, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_70, negated_conjecture, (r2_hidden(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_29]), c_0_63])).
fof(c_0_71, plain, ![X13, X14, X15]:(~l1_pre_topc(X13)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|(~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X13)))|(~v4_pre_topc(X14,X13)|~r1_tarski(X15,X14)|r1_tarski(k2_pre_topc(X13,X15),X14))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_tops_1])])])).
cnf(c_0_72, plain, (v4_pre_topc(X2,X1)|~v3_pre_topc(k3_subset_1(u1_struct_0(X1),X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_73, negated_conjecture, (v3_pre_topc(k3_subset_1(k2_pre_topc(esk1_0,esk2_0),X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_pre_topc(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_74, plain, (v2_struct_0(X1)|v3_tex_2(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_1(X2,X1)|~v2_tex_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_75, negated_conjecture, (v2_tex_2(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_29]), c_0_30]), c_0_31])]), c_0_35])).
fof(c_0_76, plain, ![X16, X17]:(((r1_tarski(X16,X17)|X16!=X17)&(r1_tarski(X17,X16)|X16!=X17))&(~r1_tarski(X16,X17)|~r1_tarski(X17,X16)|X16=X17)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_77, negated_conjecture, (r1_tarski(esk2_0,X1)|k1_zfmisc_1(u1_struct_0(esk1_0))!=k1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_78, plain, (r1_tarski(k2_pre_topc(X1,X3),X2)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v4_pre_topc(X2,X1)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_79, negated_conjecture, (v4_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_pre_topc(esk1_0,esk2_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_58]), c_0_30])]), c_0_73])).
cnf(c_0_80, negated_conjecture, (v3_tex_2(esk2_0,esk1_0)|~v1_tops_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_29]), c_0_30]), c_0_34])]), c_0_35]), c_0_75])])).
cnf(c_0_81, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_82, negated_conjecture, (r1_tarski(esk2_0,u1_struct_0(esk1_0))), inference(er,[status(thm)],[c_0_77])).
cnf(c_0_83, negated_conjecture, (r1_tarski(k2_pre_topc(esk1_0,X1),X2)|~r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_pre_topc(esk1_0,esk2_0)))|~m1_subset_1(X2,k1_zfmisc_1(k2_pre_topc(esk1_0,esk2_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_58]), c_0_30])]), c_0_79])).
cnf(c_0_84, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_pre_topc(esk1_0,esk2_0)))), inference(rw,[status(thm)],[c_0_29, c_0_58])).
cnf(c_0_85, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_86, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~v3_tex_2(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_87, negated_conjecture, (v3_tex_2(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_80, c_0_54])])).
cnf(c_0_88, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|~r1_tarski(u1_struct_0(esk1_0),esk2_0)), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_89, negated_conjecture, (r1_tarski(k2_pre_topc(esk1_0,esk2_0),X1)|~r1_tarski(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_pre_topc(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_90, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_85])).
fof(c_0_91, plain, ![X30]:~v1_subset_1(k2_subset_1(X30),X30), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc14_subset_1])])).
fof(c_0_92, plain, ![X20]:k2_subset_1(X20)=X20, inference(variable_rename,[status(thm)],[d4_subset_1])).
cnf(c_0_93, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_87])])).
cnf(c_0_94, negated_conjecture, (k2_pre_topc(esk1_0,esk2_0)=esk2_0|~r1_tarski(k2_pre_topc(esk1_0,esk2_0),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_88, c_0_58]), c_0_58])).
cnf(c_0_95, negated_conjecture, (r1_tarski(k2_pre_topc(esk1_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_84]), c_0_90])])).
cnf(c_0_96, plain, (~v1_subset_1(k2_subset_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_97, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_98, negated_conjecture, (v1_subset_1(esk2_0,k2_pre_topc(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_93, c_0_58])).
cnf(c_0_99, negated_conjecture, (k2_pre_topc(esk1_0,esk2_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_94, c_0_95])])).
cnf(c_0_100, plain, (~v1_subset_1(X1,X1)), inference(rw,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_101, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_98, c_0_99]), c_0_100]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 102
# Proof object clause steps            : 61
# Proof object formula steps           : 41
# Proof object conjectures             : 40
# Proof object clause conjectures      : 37
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 22
# Proof object simplifying inferences  : 53
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 20
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 36
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 35
# Processed clauses                    : 320
# ...of these trivial                  : 5
# ...subsumed                          : 50
# ...remaining for further processing  : 265
# Other redundant clauses eliminated   : 106
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 86
# Generated clauses                    : 1571
# ...of the previous two non-trivial   : 1422
# Contextual simplify-reflections      : 18
# Paramodulations                      : 1409
# Factorizations                       : 14
# Equation resolutions                 : 148
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 140
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 125
# Current number of unprocessed clauses: 1075
# ...number of literals in the above   : 7394
# Current number of archived formulas  : 0
# Current number of archived clauses   : 123
# Clause-clause subsumption calls (NU) : 3617
# Rec. Clause-clause subsumption calls : 1457
# Non-unit clause-clause subsumptions  : 67
# Unit Clause-clause subsumption calls : 70
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 28
# BW rewrite match successes           : 8
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 33742

# -------------------------------------------------
# User time                : 0.033 s
# System time              : 0.005 s
# Total time               : 0.037 s
# Maximum resident set size: 28788 pages
