# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/77027a55909f42e3b3b8e4a5d77e6e45.E_input', redefinition_k6_domain_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/77027a55909f42e3b3b8e4a5d77e6e45.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/77027a55909f42e3b3b8e4a5d77e6e45.E_input', t70_enumset1)).
fof(t39_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2))), file('data/MPTP/proofs/77027a55909f42e3b3b8e4a5d77e6e45.E_input', t39_yellow_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/77027a55909f42e3b3b8e4a5d77e6e45.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/77027a55909f42e3b3b8e4a5d77e6e45.E_input', dt_l1_orders_2)).
fof(t30_orders_2, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~((r1_orders_2(X1,X2,X3)&r2_orders_2(X1,X3,X2)))))), file('data/MPTP/proofs/77027a55909f42e3b3b8e4a5d77e6e45.E_input', t30_orders_2)).
fof(d10_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_orders_2(X1,X2,X3)<=>(r1_orders_2(X1,X2,X3)&~(X2=X3)))))), file('data/MPTP/proofs/77027a55909f42e3b3b8e4a5d77e6e45.E_input', d10_orders_2)).
fof(d10_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_yellow_0(X1,X2)=>(X3=k2_yellow_0(X1,X2)<=>(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))))))), file('data/MPTP/proofs/77027a55909f42e3b3b8e4a5d77e6e45.E_input', d10_yellow_0)).
fof(dt_k2_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/77027a55909f42e3b3b8e4a5d77e6e45.E_input', dt_k2_yellow_0)).
fof(t7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2)))))))), file('data/MPTP/proofs/77027a55909f42e3b3b8e4a5d77e6e45.E_input', t7_yellow_0)).
fof(t38_orders_2, axiom, ![X1]:((v3_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(~((?[X4]:((v6_orders_2(X4,X1)&m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1))))&(r2_hidden(X2,X4)&r2_hidden(X3,X4)))&(~(r1_orders_2(X1,X2,X3))&~(r1_orders_2(X1,X3,X2)))))&~(((r1_orders_2(X1,X2,X3)|r1_orders_2(X1,X3,X2))&![X4]:((v6_orders_2(X4,X1)&m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1))))=>~((r2_hidden(X2,X4)&r2_hidden(X3,X4)))))))))), file('data/MPTP/proofs/77027a55909f42e3b3b8e4a5d77e6e45.E_input', t38_orders_2)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/77027a55909f42e3b3b8e4a5d77e6e45.E_input', t4_subset)).
fof(t35_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v6_orders_2(k6_domain_1(u1_struct_0(X1),X2),X1)&m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/77027a55909f42e3b3b8e4a5d77e6e45.E_input', t35_orders_2)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/77027a55909f42e3b3b8e4a5d77e6e45.E_input', d2_tarski)).
fof(t38_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))&r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))))), file('data/MPTP/proofs/77027a55909f42e3b3b8e4a5d77e6e45.E_input', t38_yellow_0)).
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/77027a55909f42e3b3b8e4a5d77e6e45.E_input', d9_yellow_0)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/77027a55909f42e3b3b8e4a5d77e6e45.E_input', dt_k1_yellow_0)).
fof(c_0_18, plain, ![X5, X6]:(v1_xboole_0(X5)|~m1_subset_1(X6,X5)|k6_domain_1(X5,X6)=k1_tarski(X6)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_19, plain, ![X7]:k2_tarski(X7,X7)=k1_tarski(X7), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_20, plain, ![X8, X9]:k1_enumset1(X8,X8,X9)=k2_tarski(X8,X9), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_21, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2)))), inference(assume_negation,[status(cth)],[t39_yellow_0])).
cnf(c_0_22, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_25, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_21])])])])).
fof(c_0_26, plain, ![X12]:(v2_struct_0(X12)|~l1_struct_0(X12)|~v1_xboole_0(u1_struct_0(X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_27, plain, (k6_domain_1(X1,X2)=k1_enumset1(X2,X2,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_23]), c_0_24])).
cnf(c_0_28, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_29, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_30, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k1_enumset1(esk2_0,esk2_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_31, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_32, plain, ![X13]:(~l1_orders_2(X13)|l1_struct_0(X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
fof(c_0_33, plain, ![X14, X15, X16]:(~v5_orders_2(X14)|~l1_orders_2(X14)|(~m1_subset_1(X15,u1_struct_0(X14))|(~m1_subset_1(X16,u1_struct_0(X14))|(~r1_orders_2(X14,X15,X16)|~r2_orders_2(X14,X16,X15))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_orders_2])])])).
fof(c_0_34, plain, ![X17, X18, X19]:(((r1_orders_2(X17,X18,X19)|~r2_orders_2(X17,X18,X19)|~m1_subset_1(X19,u1_struct_0(X17))|~m1_subset_1(X18,u1_struct_0(X17))|~l1_orders_2(X17))&(X18!=X19|~r2_orders_2(X17,X18,X19)|~m1_subset_1(X19,u1_struct_0(X17))|~m1_subset_1(X18,u1_struct_0(X17))|~l1_orders_2(X17)))&(~r1_orders_2(X17,X18,X19)|X18=X19|r2_orders_2(X17,X18,X19)|~m1_subset_1(X19,u1_struct_0(X17))|~m1_subset_1(X18,u1_struct_0(X17))|~l1_orders_2(X17))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_orders_2])])])])])).
fof(c_0_35, plain, ![X20, X21, X22, X23]:(((r1_lattice3(X20,X21,X22)|X22!=k2_yellow_0(X20,X21)|~r2_yellow_0(X20,X21)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))&(~m1_subset_1(X23,u1_struct_0(X20))|(~r1_lattice3(X20,X21,X23)|r1_orders_2(X20,X23,X22))|X22!=k2_yellow_0(X20,X21)|~r2_yellow_0(X20,X21)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20)))&((m1_subset_1(esk3_3(X20,X21,X22),u1_struct_0(X20))|~r1_lattice3(X20,X21,X22)|X22=k2_yellow_0(X20,X21)|~r2_yellow_0(X20,X21)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))&((r1_lattice3(X20,X21,esk3_3(X20,X21,X22))|~r1_lattice3(X20,X21,X22)|X22=k2_yellow_0(X20,X21)|~r2_yellow_0(X20,X21)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))&(~r1_orders_2(X20,esk3_3(X20,X21,X22),X22)|~r1_lattice3(X20,X21,X22)|X22=k2_yellow_0(X20,X21)|~r2_yellow_0(X20,X21)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_yellow_0])])])])])).
fof(c_0_36, plain, ![X25, X26]:(~l1_orders_2(X25)|m1_subset_1(k2_yellow_0(X25,X26),u1_struct_0(X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_yellow_0])])).
fof(c_0_37, plain, ![X27, X28, X29]:((~r1_lattice3(X27,k1_tarski(X29),X28)|r1_orders_2(X27,X28,X29)|~m1_subset_1(X29,u1_struct_0(X27))|~m1_subset_1(X28,u1_struct_0(X27))|~l1_orders_2(X27))&((~r1_orders_2(X27,X28,X29)|r1_lattice3(X27,k1_tarski(X29),X28)|~m1_subset_1(X29,u1_struct_0(X27))|~m1_subset_1(X28,u1_struct_0(X27))|~l1_orders_2(X27))&((~r2_lattice3(X27,k1_tarski(X29),X28)|r1_orders_2(X27,X29,X28)|~m1_subset_1(X29,u1_struct_0(X27))|~m1_subset_1(X28,u1_struct_0(X27))|~l1_orders_2(X27))&(~r1_orders_2(X27,X29,X28)|r2_lattice3(X27,k1_tarski(X29),X28)|~m1_subset_1(X29,u1_struct_0(X27))|~m1_subset_1(X28,u1_struct_0(X27))|~l1_orders_2(X27))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])).
fof(c_0_38, plain, ![X30, X31, X32, X33]:((~v6_orders_2(X33,X30)|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X30)))|(~r2_hidden(X31,X33)|~r2_hidden(X32,X33))|(r1_orders_2(X30,X31,X32)|r1_orders_2(X30,X32,X31))|~m1_subset_1(X32,u1_struct_0(X30))|~m1_subset_1(X31,u1_struct_0(X30))|(~v3_orders_2(X30)|~l1_orders_2(X30)))&((((v6_orders_2(esk4_3(X30,X31,X32),X30)|~r1_orders_2(X30,X31,X32)|~m1_subset_1(X32,u1_struct_0(X30))|~m1_subset_1(X31,u1_struct_0(X30))|(~v3_orders_2(X30)|~l1_orders_2(X30)))&(m1_subset_1(esk4_3(X30,X31,X32),k1_zfmisc_1(u1_struct_0(X30)))|~r1_orders_2(X30,X31,X32)|~m1_subset_1(X32,u1_struct_0(X30))|~m1_subset_1(X31,u1_struct_0(X30))|(~v3_orders_2(X30)|~l1_orders_2(X30))))&((r2_hidden(X31,esk4_3(X30,X31,X32))|~r1_orders_2(X30,X31,X32)|~m1_subset_1(X32,u1_struct_0(X30))|~m1_subset_1(X31,u1_struct_0(X30))|(~v3_orders_2(X30)|~l1_orders_2(X30)))&(r2_hidden(X32,esk4_3(X30,X31,X32))|~r1_orders_2(X30,X31,X32)|~m1_subset_1(X32,u1_struct_0(X30))|~m1_subset_1(X31,u1_struct_0(X30))|(~v3_orders_2(X30)|~l1_orders_2(X30)))))&(((v6_orders_2(esk4_3(X30,X31,X32),X30)|~r1_orders_2(X30,X32,X31)|~m1_subset_1(X32,u1_struct_0(X30))|~m1_subset_1(X31,u1_struct_0(X30))|(~v3_orders_2(X30)|~l1_orders_2(X30)))&(m1_subset_1(esk4_3(X30,X31,X32),k1_zfmisc_1(u1_struct_0(X30)))|~r1_orders_2(X30,X32,X31)|~m1_subset_1(X32,u1_struct_0(X30))|~m1_subset_1(X31,u1_struct_0(X30))|(~v3_orders_2(X30)|~l1_orders_2(X30))))&((r2_hidden(X31,esk4_3(X30,X31,X32))|~r1_orders_2(X30,X32,X31)|~m1_subset_1(X32,u1_struct_0(X30))|~m1_subset_1(X31,u1_struct_0(X30))|(~v3_orders_2(X30)|~l1_orders_2(X30)))&(r2_hidden(X32,esk4_3(X30,X31,X32))|~r1_orders_2(X30,X32,X31)|~m1_subset_1(X32,u1_struct_0(X30))|~m1_subset_1(X31,u1_struct_0(X30))|(~v3_orders_2(X30)|~l1_orders_2(X30))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t38_orders_2])])])])])])).
fof(c_0_39, plain, ![X35, X36, X37]:(~r2_hidden(X35,X36)|~m1_subset_1(X36,k1_zfmisc_1(X37))|m1_subset_1(X35,X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_40, plain, ![X38, X39]:((v6_orders_2(k6_domain_1(u1_struct_0(X38),X39),X38)|~m1_subset_1(X39,u1_struct_0(X38))|(v2_struct_0(X38)|(~v3_orders_2(X38)|~l1_orders_2(X38))))&(m1_subset_1(k6_domain_1(u1_struct_0(X38),X39),k1_zfmisc_1(u1_struct_0(X38)))|~m1_subset_1(X39,u1_struct_0(X38))|(v2_struct_0(X38)|(~v3_orders_2(X38)|~l1_orders_2(X38))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t35_orders_2])])])])])).
cnf(c_0_41, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k1_enumset1(esk2_0,esk2_0,esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])).
cnf(c_0_42, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_44, plain, ![X40, X41, X42, X43, X44, X45, X46, X47]:(((~r2_hidden(X43,X42)|(X43=X40|X43=X41)|X42!=k2_tarski(X40,X41))&((X44!=X40|r2_hidden(X44,X42)|X42!=k2_tarski(X40,X41))&(X44!=X41|r2_hidden(X44,X42)|X42!=k2_tarski(X40,X41))))&(((esk5_3(X45,X46,X47)!=X45|~r2_hidden(esk5_3(X45,X46,X47),X47)|X47=k2_tarski(X45,X46))&(esk5_3(X45,X46,X47)!=X46|~r2_hidden(esk5_3(X45,X46,X47),X47)|X47=k2_tarski(X45,X46)))&(r2_hidden(esk5_3(X45,X46,X47),X47)|(esk5_3(X45,X46,X47)=X45|esk5_3(X45,X46,X47)=X46)|X47=k2_tarski(X45,X46)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_45, plain, (~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)|~r2_orders_2(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_46, plain, (X2=X3|r2_orders_2(X1,X2,X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_47, plain, (r1_orders_2(X2,X1,X4)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|X4!=k2_yellow_0(X2,X3)|~r2_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_48, plain, (m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_49, plain, (r1_orders_2(X1,X3,X2)|~r1_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_50, plain, (r1_lattice3(X1,X2,X3)|X3!=k2_yellow_0(X1,X2)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_51, plain, (r1_orders_2(X2,X3,X4)|r1_orders_2(X2,X4,X3)|~v6_orders_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X3,X1)|~r2_hidden(X4,X1)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~v3_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_52, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_53, plain, (m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_54, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k1_enumset1(esk2_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43])])).
cnf(c_0_55, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_56, plain, (v6_orders_2(k6_domain_1(u1_struct_0(X1),X2),X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_57, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_58, plain, (X1=X2|~r1_orders_2(X3,X2,X1)|~r1_orders_2(X3,X1,X2)|~l1_orders_2(X3)|~v5_orders_2(X3)|~m1_subset_1(X1,u1_struct_0(X3))|~m1_subset_1(X2,u1_struct_0(X3))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_59, plain, (r1_orders_2(X1,X2,k2_yellow_0(X1,X3))|~r1_lattice3(X1,X3,X2)|~r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_47]), c_0_48])).
cnf(c_0_60, plain, (r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,k1_enumset1(X2,X2,X2),X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_23]), c_0_24])).
cnf(c_0_61, plain, (r1_lattice3(X1,X2,k2_yellow_0(X1,X2))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_50]), c_0_48])).
cnf(c_0_62, plain, (r1_orders_2(X1,X2,X3)|r1_orders_2(X1,X3,X2)|~r2_hidden(X3,X4)|~r2_hidden(X2,X4)|~v6_orders_2(X4,X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_51, c_0_52]), c_0_52])).
cnf(c_0_63, negated_conjecture, (m1_subset_1(k1_enumset1(esk2_0,esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_43]), c_0_55]), c_0_28])]), c_0_31])).
cnf(c_0_64, negated_conjecture, (v6_orders_2(k1_enumset1(esk2_0,esk2_0,esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_54]), c_0_43]), c_0_55]), c_0_28])]), c_0_31])).
cnf(c_0_65, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X4,X2)), inference(rw,[status(thm)],[c_0_57, c_0_24])).
cnf(c_0_66, plain, (k2_yellow_0(X1,X2)=X3|~r1_lattice3(X1,X2,X3)|~r2_yellow_0(X1,X2)|~r1_orders_2(X1,k2_yellow_0(X1,X2),X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_48])).
cnf(c_0_67, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_enumset1(X2,X2,X2)),X2)|~r2_yellow_0(X1,k1_enumset1(X2,X2,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_48])).
cnf(c_0_68, plain, (r1_lattice3(X1,k1_tarski(X3),X2)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_69, plain, ![X49, X50]:((r1_yellow_0(X49,k6_domain_1(u1_struct_0(X49),X50))|~m1_subset_1(X50,u1_struct_0(X49))|(v2_struct_0(X49)|(~v3_orders_2(X49)|(~v5_orders_2(X49)|~l1_orders_2(X49)))))&(r2_yellow_0(X49,k6_domain_1(u1_struct_0(X49),X50))|~m1_subset_1(X50,u1_struct_0(X49))|(v2_struct_0(X49)|(~v3_orders_2(X49)|(~v5_orders_2(X49)|~l1_orders_2(X49)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t38_yellow_0])])])])])).
cnf(c_0_70, negated_conjecture, (r1_orders_2(esk1_0,X1,X2)|r1_orders_2(esk1_0,X2,X1)|~r2_hidden(X2,k1_enumset1(esk2_0,esk2_0,esk2_0))|~r2_hidden(X1,k1_enumset1(esk2_0,esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64]), c_0_43]), c_0_55])])).
cnf(c_0_71, plain, (r2_hidden(X1,k1_enumset1(X2,X2,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_65])])).
cnf(c_0_72, plain, (k2_yellow_0(X1,k1_enumset1(X2,X2,X2))=X2|~r1_lattice3(X1,k1_enumset1(X2,X2,X2),X2)|~r2_yellow_0(X1,k1_enumset1(X2,X2,X2))|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_73, plain, (r1_lattice3(X1,k1_enumset1(X3,X3,X3),X2)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_23]), c_0_24])).
cnf(c_0_74, plain, (r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_75, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_76, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,X1)|r1_orders_2(esk1_0,X1,esk2_0)|~r2_hidden(X1,k1_enumset1(esk2_0,esk2_0,esk2_0))), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_77, plain, (k2_yellow_0(X1,k1_enumset1(X2,X2,X2))=X2|~r2_yellow_0(X1,k1_enumset1(X2,X2,X2))|~r1_orders_2(X1,X2,X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_78, negated_conjecture, (r2_yellow_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_54]), c_0_43]), c_0_75]), c_0_55]), c_0_28])]), c_0_31])).
cnf(c_0_79, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_76, c_0_71])).
fof(c_0_80, plain, ![X51, X52, X53, X54]:(((r2_lattice3(X51,X52,X53)|X53!=k1_yellow_0(X51,X52)|~r1_yellow_0(X51,X52)|~m1_subset_1(X53,u1_struct_0(X51))|~l1_orders_2(X51))&(~m1_subset_1(X54,u1_struct_0(X51))|(~r2_lattice3(X51,X52,X54)|r1_orders_2(X51,X53,X54))|X53!=k1_yellow_0(X51,X52)|~r1_yellow_0(X51,X52)|~m1_subset_1(X53,u1_struct_0(X51))|~l1_orders_2(X51)))&((m1_subset_1(esk6_3(X51,X52,X53),u1_struct_0(X51))|~r2_lattice3(X51,X52,X53)|X53=k1_yellow_0(X51,X52)|~r1_yellow_0(X51,X52)|~m1_subset_1(X53,u1_struct_0(X51))|~l1_orders_2(X51))&((r2_lattice3(X51,X52,esk6_3(X51,X52,X53))|~r2_lattice3(X51,X52,X53)|X53=k1_yellow_0(X51,X52)|~r1_yellow_0(X51,X52)|~m1_subset_1(X53,u1_struct_0(X51))|~l1_orders_2(X51))&(~r1_orders_2(X51,X53,esk6_3(X51,X52,X53))|~r2_lattice3(X51,X52,X53)|X53=k1_yellow_0(X51,X52)|~r1_yellow_0(X51,X52)|~m1_subset_1(X53,u1_struct_0(X51))|~l1_orders_2(X51))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
fof(c_0_81, plain, ![X56, X57]:(~l1_orders_2(X56)|m1_subset_1(k1_yellow_0(X56,X57),u1_struct_0(X56))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
cnf(c_0_82, negated_conjecture, (k2_yellow_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk2_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_79]), c_0_43]), c_0_75]), c_0_28])])).
cnf(c_0_83, plain, (r1_orders_2(X1,X2,X3)|~r2_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_84, plain, (r2_lattice3(X1,X2,X3)|X3!=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_85, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_86, negated_conjecture, (k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_87, negated_conjecture, (esk2_0=X1|~r1_lattice3(esk1_0,k1_enumset1(esk2_0,esk2_0,esk2_0),X1)|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_82]), c_0_78]), c_0_43]), c_0_75])])).
cnf(c_0_88, plain, (r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r2_lattice3(X1,k1_enumset1(X2,X2,X2),X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_83, c_0_23]), c_0_24])).
cnf(c_0_89, plain, (r2_lattice3(X1,X2,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_84]), c_0_85])).
cnf(c_0_90, plain, (r1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_91, negated_conjecture, (k1_yellow_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk2_0))!=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_54]), c_0_54])).
cnf(c_0_92, negated_conjecture, (esk2_0=X1|~r1_orders_2(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_73]), c_0_43]), c_0_28])])).
cnf(c_0_93, plain, (r1_orders_2(X1,X2,k1_yellow_0(X1,k1_enumset1(X2,X2,X2)))|~r1_yellow_0(X1,k1_enumset1(X2,X2,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_85])).
cnf(c_0_94, negated_conjecture, (r1_yellow_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_54]), c_0_43]), c_0_75]), c_0_55]), c_0_28])]), c_0_31])).
cnf(c_0_95, negated_conjecture, (k1_yellow_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk2_0))!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_82])])).
cnf(c_0_96, plain, (r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~r1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_97, negated_conjecture, (~r1_orders_2(esk1_0,k1_yellow_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk2_0)),esk2_0)|~m1_subset_1(k1_yellow_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk2_0)),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_94]), c_0_43]), c_0_28])]), c_0_95])).
cnf(c_0_98, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),X3)|~r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_96]), c_0_85])).
cnf(c_0_99, negated_conjecture, (~r2_lattice3(esk1_0,k1_enumset1(esk2_0,esk2_0,esk2_0),esk2_0)|~m1_subset_1(k1_yellow_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk2_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_94]), c_0_43]), c_0_28])])).
cnf(c_0_100, plain, (r2_lattice3(X1,k1_tarski(X2),X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_101, negated_conjecture, (~r2_lattice3(esk1_0,k1_enumset1(esk2_0,esk2_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_85]), c_0_43])])).
cnf(c_0_102, plain, (r2_lattice3(X1,k1_enumset1(X2,X2,X2),X3)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_100, c_0_23]), c_0_24])).
cnf(c_0_103, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_79]), c_0_43]), c_0_28])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 104
# Proof object clause steps            : 67
# Proof object formula steps           : 37
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 30
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 23
# Proof object simplifying inferences  : 86
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 51
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 49
# Processed clauses                    : 175
# ...of these trivial                  : 1
# ...subsumed                          : 10
# ...remaining for further processing  : 164
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 4
# Generated clauses                    : 373
# ...of the previous two non-trivial   : 342
# Contextual simplify-reflections      : 16
# Paramodulations                      : 363
# Factorizations                       : 0
# Equation resolutions                 : 12
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 103
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 86
# Current number of unprocessed clauses: 264
# ...number of literals in the above   : 2048
# Current number of archived formulas  : 0
# Current number of archived clauses   : 55
# Clause-clause subsumption calls (NU) : 2960
# Rec. Clause-clause subsumption calls : 163
# Non-unit clause-clause subsumptions  : 26
# Unit Clause-clause subsumption calls : 85
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 15073

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.005 s
# Total time               : 0.021 s
# Maximum resident set size: 28708 pages
