# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.007 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d15_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k3_waybel_0(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>?[X5]:(m1_subset_1(X5,u1_struct_0(X1))&(r1_orders_2(X1,X4,X5)&r2_hidden(X5,X2))))))))), file('data/MPTP/proofs/77069c138a6b402d9939ecfcef2667f1.E_input', d15_waybel_0)).
fof(d17_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>k5_waybel_0(X1,X2)=k3_waybel_0(X1,k6_domain_1(u1_struct_0(X1),X2)))), file('data/MPTP/proofs/77069c138a6b402d9939ecfcef2667f1.E_input', d17_waybel_0)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/77069c138a6b402d9939ecfcef2667f1.E_input', redefinition_k6_domain_1)).
fof(dt_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))), file('data/MPTP/proofs/77069c138a6b402d9939ecfcef2667f1.E_input', dt_k6_domain_1)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/77069c138a6b402d9939ecfcef2667f1.E_input', t5_subset)).
fof(dt_k5_waybel_0, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/77069c138a6b402d9939ecfcef2667f1.E_input', dt_k5_waybel_0)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/77069c138a6b402d9939ecfcef2667f1.E_input', d1_tarski)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/77069c138a6b402d9939ecfcef2667f1.E_input', d1_xboole_0)).
fof(t17_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2))))), file('data/MPTP/proofs/77069c138a6b402d9939ecfcef2667f1.E_input', t17_waybel_0)).
fof(d9_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)<=>r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1)))))), file('data/MPTP/proofs/77069c138a6b402d9939ecfcef2667f1.E_input', d9_orders_2)).
fof(cc4_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/77069c138a6b402d9939ecfcef2667f1.E_input', cc4_relset_1)).
fof(dt_u1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))), file('data/MPTP/proofs/77069c138a6b402d9939ecfcef2667f1.E_input', dt_u1_orders_2)).
fof(c_0_12, plain, ![X24, X25, X26, X27, X29, X31]:((((m1_subset_1(esk3_4(X24,X25,X26,X27),u1_struct_0(X24))|~r2_hidden(X27,X26)|~m1_subset_1(X27,u1_struct_0(X24))|X26!=k3_waybel_0(X24,X25)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X24)))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_orders_2(X24))&((r1_orders_2(X24,X27,esk3_4(X24,X25,X26,X27))|~r2_hidden(X27,X26)|~m1_subset_1(X27,u1_struct_0(X24))|X26!=k3_waybel_0(X24,X25)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X24)))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_orders_2(X24))&(r2_hidden(esk3_4(X24,X25,X26,X27),X25)|~r2_hidden(X27,X26)|~m1_subset_1(X27,u1_struct_0(X24))|X26!=k3_waybel_0(X24,X25)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X24)))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_orders_2(X24))))&(~m1_subset_1(X29,u1_struct_0(X24))|(~r1_orders_2(X24,X27,X29)|~r2_hidden(X29,X25))|r2_hidden(X27,X26)|~m1_subset_1(X27,u1_struct_0(X24))|X26!=k3_waybel_0(X24,X25)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X24)))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_orders_2(X24)))&((m1_subset_1(esk4_3(X24,X25,X26),u1_struct_0(X24))|X26=k3_waybel_0(X24,X25)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X24)))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_orders_2(X24))&((~r2_hidden(esk4_3(X24,X25,X26),X26)|(~m1_subset_1(X31,u1_struct_0(X24))|(~r1_orders_2(X24,esk4_3(X24,X25,X26),X31)|~r2_hidden(X31,X25)))|X26=k3_waybel_0(X24,X25)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X24)))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_orders_2(X24))&((m1_subset_1(esk5_3(X24,X25,X26),u1_struct_0(X24))|r2_hidden(esk4_3(X24,X25,X26),X26)|X26=k3_waybel_0(X24,X25)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X24)))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_orders_2(X24))&((r1_orders_2(X24,esk4_3(X24,X25,X26),esk5_3(X24,X25,X26))|r2_hidden(esk4_3(X24,X25,X26),X26)|X26=k3_waybel_0(X24,X25)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X24)))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_orders_2(X24))&(r2_hidden(esk5_3(X24,X25,X26),X25)|r2_hidden(esk4_3(X24,X25,X26),X26)|X26=k3_waybel_0(X24,X25)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X24)))|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))|~l1_orders_2(X24))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d15_waybel_0])])])])])).
fof(c_0_13, plain, ![X33, X34]:(v2_struct_0(X33)|~l1_orders_2(X33)|(~m1_subset_1(X34,u1_struct_0(X33))|k5_waybel_0(X33,X34)=k3_waybel_0(X33,k6_domain_1(u1_struct_0(X33),X34)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_waybel_0])])])])).
fof(c_0_14, plain, ![X35, X36]:(v1_xboole_0(X35)|~m1_subset_1(X36,X35)|k6_domain_1(X35,X36)=k1_tarski(X36)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_15, plain, ![X37, X38]:(v1_xboole_0(X37)|~m1_subset_1(X38,X37)|m1_subset_1(k6_domain_1(X37,X38),k1_zfmisc_1(X37))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_domain_1])])])).
fof(c_0_16, plain, ![X39, X40, X41]:(~r2_hidden(X39,X40)|(~m1_subset_1(X40,k1_zfmisc_1(X41))|~v1_xboole_0(X41))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_17, plain, ![X42, X43]:(v2_struct_0(X42)|~l1_orders_2(X42)|~m1_subset_1(X43,u1_struct_0(X42))|m1_subset_1(k5_waybel_0(X42,X43),k1_zfmisc_1(u1_struct_0(X42)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_waybel_0])])])).
cnf(c_0_18, plain, (r2_hidden(X3,X5)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,X3,X1)|~r2_hidden(X1,X4)|~m1_subset_1(X3,u1_struct_0(X2))|X5!=k3_waybel_0(X2,X4)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_19, plain, ![X17, X18, X19, X20, X21, X22]:(((~r2_hidden(X19,X18)|X19=X17|X18!=k1_tarski(X17))&(X20!=X17|r2_hidden(X20,X18)|X18!=k1_tarski(X17)))&((~r2_hidden(esk2_2(X21,X22),X22)|esk2_2(X21,X22)!=X21|X22=k1_tarski(X21))&(r2_hidden(esk2_2(X21,X22),X22)|esk2_2(X21,X22)=X21|X22=k1_tarski(X21)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_20, plain, (r2_hidden(esk3_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, plain, (v2_struct_0(X1)|k5_waybel_0(X1,X2)=k3_waybel_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23, plain, (v1_xboole_0(X1)|m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, plain, (v2_struct_0(X1)|m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (r2_hidden(X1,k3_waybel_0(X2,X3))|~r1_orders_2(X2,X1,X4)|~m1_subset_1(k3_waybel_0(X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)|~r2_hidden(X4,X3)), inference(er,[status(thm)],[c_0_18])).
fof(c_0_27, plain, ![X6, X7, X8]:((~v1_xboole_0(X6)|~r2_hidden(X7,X6))&(r2_hidden(esk1_1(X8),X8)|v1_xboole_0(X8))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_28, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (r1_orders_2(X1,X2,esk3_4(X1,X3,X4,X2))|~r2_hidden(X2,X4)|~m1_subset_1(X2,u1_struct_0(X1))|X4!=k3_waybel_0(X1,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_30, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_31, plain, (r2_hidden(esk3_4(X1,X2,k3_waybel_0(X1,X2),X3),X2)|~m1_subset_1(k3_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(X3,k3_waybel_0(X1,X2))), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_32, plain, (k3_waybel_0(X1,k1_tarski(X2))=k5_waybel_0(X1,X2)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_33, plain, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(spm,[status(thm)],[c_0_23, c_0_22])).
cnf(c_0_34, plain, (v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(X3,k5_waybel_0(X1,X2))|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_35, plain, (v2_struct_0(X1)|r2_hidden(X2,k5_waybel_0(X1,X3))|~r1_orders_2(X1,X2,X4)|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(X4,k6_domain_1(u1_struct_0(X1),X3))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_21]), c_0_25])).
cnf(c_0_36, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_28])])).
cnf(c_0_38, plain, (r1_orders_2(X1,X2,esk3_4(X1,X3,k3_waybel_0(X1,X3),X2))|~m1_subset_1(k3_waybel_0(X1,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(X2,k3_waybel_0(X1,X3))), inference(er,[status(thm)],[c_0_29])).
cnf(c_0_39, plain, (X1=X2|~r2_hidden(X1,k1_tarski(X2))), inference(er,[status(thm)],[c_0_30])).
cnf(c_0_40, plain, (v2_struct_0(X1)|r2_hidden(esk3_4(X1,k1_tarski(X2),k5_waybel_0(X1,X2),X3),k1_tarski(X2))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(X3,k5_waybel_0(X1,X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33]), c_0_25]), c_0_34])).
fof(c_0_41, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2)))))), inference(assume_negation,[status(cth)],[t17_waybel_0])).
cnf(c_0_42, plain, (v2_struct_0(X1)|r2_hidden(X2,k5_waybel_0(X1,X3))|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,X2,X4)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(X4,k6_domain_1(u1_struct_0(X1),X3))), inference(spm,[status(thm)],[c_0_35, c_0_23])).
cnf(c_0_43, plain, (r2_hidden(esk1_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_44, plain, (~v1_xboole_0(k1_tarski(X1))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
fof(c_0_45, plain, ![X10, X11, X12]:((~r1_orders_2(X10,X11,X12)|r2_hidden(k4_tarski(X11,X12),u1_orders_2(X10))|~m1_subset_1(X12,u1_struct_0(X10))|~m1_subset_1(X11,u1_struct_0(X10))|~l1_orders_2(X10))&(~r2_hidden(k4_tarski(X11,X12),u1_orders_2(X10))|r1_orders_2(X10,X11,X12)|~m1_subset_1(X12,u1_struct_0(X10))|~m1_subset_1(X11,u1_struct_0(X10))|~l1_orders_2(X10))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_orders_2])])])])).
cnf(c_0_46, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,esk3_4(X1,k1_tarski(X3),k5_waybel_0(X1,X3),X2))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(X2,k5_waybel_0(X1,X3))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_32]), c_0_33]), c_0_25]), c_0_34])).
cnf(c_0_47, plain, (esk3_4(X1,k1_tarski(X2),k5_waybel_0(X1,X2),X3)=X2|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(X3,k5_waybel_0(X1,X2))), inference(spm,[status(thm)],[c_0_39, c_0_40])).
fof(c_0_48, negated_conjecture, ((~v2_struct_0(esk6_0)&l1_orders_2(esk6_0))&(m1_subset_1(esk7_0,u1_struct_0(esk6_0))&(m1_subset_1(esk8_0,u1_struct_0(esk6_0))&((~r2_hidden(esk8_0,k5_waybel_0(esk6_0,esk7_0))|~r1_orders_2(esk6_0,esk8_0,esk7_0))&(r2_hidden(esk8_0,k5_waybel_0(esk6_0,esk7_0))|r1_orders_2(esk6_0,esk8_0,esk7_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_41])])])])).
cnf(c_0_49, plain, (v2_struct_0(X1)|r2_hidden(X2,k5_waybel_0(X1,X3))|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,X2,X4)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(X4,k1_tarski(X3))), inference(spm,[status(thm)],[c_0_42, c_0_22])).
cnf(c_0_50, plain, (esk1_1(k1_tarski(X1))=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_43]), c_0_44])).
cnf(c_0_51, plain, (r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_52, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(X2,k5_waybel_0(X1,X3))), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_53, negated_conjecture, (r2_hidden(esk8_0,k5_waybel_0(esk6_0,esk7_0))|r1_orders_2(esk6_0,esk8_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(esk8_0,u1_struct_0(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_55, negated_conjecture, (m1_subset_1(esk7_0,u1_struct_0(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_56, negated_conjecture, (l1_orders_2(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_57, negated_conjecture, (~v2_struct_0(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_58, plain, ![X13, X14, X15]:(~v1_xboole_0(X13)|(~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X14,X13)))|v1_xboole_0(X15))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])).
fof(c_0_59, plain, ![X16]:(~l1_orders_2(X16)|m1_subset_1(u1_orders_2(X16),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X16),u1_struct_0(X16))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_orders_2])])).
cnf(c_0_60, negated_conjecture, (~r2_hidden(esk8_0,k5_waybel_0(esk6_0,esk7_0))|~r1_orders_2(esk6_0,esk8_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_61, plain, (v2_struct_0(X1)|r2_hidden(X2,k5_waybel_0(X1,X3))|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_43]), c_0_50]), c_0_50])]), c_0_44])).
cnf(c_0_62, plain, (~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_xboole_0(u1_orders_2(X1))), inference(spm,[status(thm)],[c_0_36, c_0_51])).
cnf(c_0_63, negated_conjecture, (r1_orders_2(esk6_0,esk8_0,esk7_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54]), c_0_55]), c_0_56])]), c_0_57])).
cnf(c_0_64, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_65, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_66, negated_conjecture, (v1_xboole_0(u1_struct_0(esk6_0))|~r1_orders_2(esk6_0,esk8_0,esk7_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_54]), c_0_55]), c_0_56])]), c_0_57])).
cnf(c_0_67, negated_conjecture, (~v1_xboole_0(u1_orders_2(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_55]), c_0_54]), c_0_56])])).
cnf(c_0_68, plain, (v1_xboole_0(u1_orders_2(X1))|~l1_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_69, negated_conjecture, (v1_xboole_0(u1_struct_0(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_63])])).
cnf(c_0_70, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_56]), c_0_69])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 71
# Proof object clause steps            : 46
# Proof object formula steps           : 25
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 37
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 509
# ...of these trivial                  : 0
# ...subsumed                          : 315
# ...remaining for further processing  : 194
# Other redundant clauses eliminated   : 11
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 21
# Backward-rewritten                   : 41
# Generated clauses                    : 925
# ...of the previous two non-trivial   : 890
# Contextual simplify-reflections      : 24
# Paramodulations                      : 913
# Factorizations                       : 2
# Equation resolutions                 : 11
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 96
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 85
# Current number of unprocessed clauses: 400
# ...number of literals in the above   : 2793
# Current number of archived formulas  : 0
# Current number of archived clauses   : 92
# Clause-clause subsumption calls (NU) : 9859
# Rec. Clause-clause subsumption calls : 1285
# Non-unit clause-clause subsumptions  : 355
# Unit Clause-clause subsumption calls : 103
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 32113

# -------------------------------------------------
# User time                : 0.027 s
# System time              : 0.000 s
# Total time               : 0.027 s
# Maximum resident set size: 28660 pages
