# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d8_yellow_6, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(l1_waybel_0(X2,X1)=>![X3]:(l1_waybel_0(X3,X1)=>(m1_yellow_6(X3,X1,X2)<=>(m1_yellow_0(X3,X2)&u1_waybel_0(X1,X3)=k2_partfun1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2),u1_struct_0(X3))))))), file('data/MPTP/proofs/771e2cc7ddbd46fba10c1bd2a144fd74.E_input', d8_yellow_6)).
fof(dt_m1_yellow_6, axiom, ![X1, X2]:((l1_struct_0(X1)&l1_waybel_0(X2,X1))=>![X3]:(m1_yellow_6(X3,X1,X2)=>l1_waybel_0(X3,X1))), file('data/MPTP/proofs/771e2cc7ddbd46fba10c1bd2a144fd74.E_input', dt_m1_yellow_6)).
fof(t20_yellow_6, conjecture, ![X1]:(l1_struct_0(X1)=>![X2]:(l1_waybel_0(X2,X1)=>![X3]:(m1_yellow_6(X3,X1,X2)=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>![X6]:(m1_subset_1(X6,u1_struct_0(X3))=>![X7]:(m1_subset_1(X7,u1_struct_0(X3))=>((X4=X6&(X5=X7&r1_orders_2(X3,X6,X7)))=>r1_orders_2(X2,X4,X5))))))))), file('data/MPTP/proofs/771e2cc7ddbd46fba10c1bd2a144fd74.E_input', t20_yellow_6)).
fof(d13_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(m1_yellow_0(X2,X1)<=>(r1_tarski(u1_struct_0(X2),u1_struct_0(X1))&r1_tarski(u1_orders_2(X2),u1_orders_2(X1)))))), file('data/MPTP/proofs/771e2cc7ddbd46fba10c1bd2a144fd74.E_input', d13_yellow_0)).
fof(dt_m1_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_yellow_0(X2,X1)=>l1_orders_2(X2))), file('data/MPTP/proofs/771e2cc7ddbd46fba10c1bd2a144fd74.E_input', dt_m1_yellow_0)).
fof(dt_l1_waybel_0, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(l1_waybel_0(X2,X1)=>l1_orders_2(X2))), file('data/MPTP/proofs/771e2cc7ddbd46fba10c1bd2a144fd74.E_input', dt_l1_waybel_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/771e2cc7ddbd46fba10c1bd2a144fd74.E_input', t3_subset)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/771e2cc7ddbd46fba10c1bd2a144fd74.E_input', t2_subset)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/771e2cc7ddbd46fba10c1bd2a144fd74.E_input', fc1_subset_1)).
fof(d9_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)<=>r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1)))))), file('data/MPTP/proofs/771e2cc7ddbd46fba10c1bd2a144fd74.E_input', d9_orders_2)).
fof(d4_tarski, axiom, ![X1, X2]:(X2=k3_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X3,X4)&r2_hidden(X4,X1)))), file('data/MPTP/proofs/771e2cc7ddbd46fba10c1bd2a144fd74.E_input', d4_tarski)).
fof(t99_zfmisc_1, axiom, ![X1]:k3_tarski(k1_zfmisc_1(X1))=X1, file('data/MPTP/proofs/771e2cc7ddbd46fba10c1bd2a144fd74.E_input', t99_zfmisc_1)).
fof(c_0_12, plain, ![X25, X26, X27]:(((m1_yellow_0(X27,X26)|~m1_yellow_6(X27,X25,X26)|~l1_waybel_0(X27,X25)|~l1_waybel_0(X26,X25)|~l1_struct_0(X25))&(u1_waybel_0(X25,X27)=k2_partfun1(u1_struct_0(X26),u1_struct_0(X25),u1_waybel_0(X25,X26),u1_struct_0(X27))|~m1_yellow_6(X27,X25,X26)|~l1_waybel_0(X27,X25)|~l1_waybel_0(X26,X25)|~l1_struct_0(X25)))&(~m1_yellow_0(X27,X26)|u1_waybel_0(X25,X27)!=k2_partfun1(u1_struct_0(X26),u1_struct_0(X25),u1_waybel_0(X25,X26),u1_struct_0(X27))|m1_yellow_6(X27,X25,X26)|~l1_waybel_0(X27,X25)|~l1_waybel_0(X26,X25)|~l1_struct_0(X25))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_yellow_6])])])])).
fof(c_0_13, plain, ![X31, X32, X33]:(~l1_struct_0(X31)|~l1_waybel_0(X32,X31)|(~m1_yellow_6(X33,X31,X32)|l1_waybel_0(X33,X31))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_yellow_6])])])).
fof(c_0_14, negated_conjecture, ~(![X1]:(l1_struct_0(X1)=>![X2]:(l1_waybel_0(X2,X1)=>![X3]:(m1_yellow_6(X3,X1,X2)=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>![X6]:(m1_subset_1(X6,u1_struct_0(X3))=>![X7]:(m1_subset_1(X7,u1_struct_0(X3))=>((X4=X6&(X5=X7&r1_orders_2(X3,X6,X7)))=>r1_orders_2(X2,X4,X5)))))))))), inference(assume_negation,[status(cth)],[t20_yellow_6])).
fof(c_0_15, plain, ![X28, X29]:(((r1_tarski(u1_struct_0(X29),u1_struct_0(X28))|~m1_yellow_0(X29,X28)|~l1_orders_2(X29)|~l1_orders_2(X28))&(r1_tarski(u1_orders_2(X29),u1_orders_2(X28))|~m1_yellow_0(X29,X28)|~l1_orders_2(X29)|~l1_orders_2(X28)))&(~r1_tarski(u1_struct_0(X29),u1_struct_0(X28))|~r1_tarski(u1_orders_2(X29),u1_orders_2(X28))|m1_yellow_0(X29,X28)|~l1_orders_2(X29)|~l1_orders_2(X28))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d13_yellow_0])])])])).
fof(c_0_16, plain, ![X23, X24]:(~l1_orders_2(X23)|(~m1_yellow_0(X24,X23)|l1_orders_2(X24))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_yellow_0])])])).
cnf(c_0_17, plain, (m1_yellow_0(X1,X2)|~m1_yellow_6(X1,X3,X2)|~l1_waybel_0(X1,X3)|~l1_waybel_0(X2,X3)|~l1_struct_0(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, plain, (l1_waybel_0(X3,X1)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_yellow_6(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_19, negated_conjecture, (l1_struct_0(esk1_0)&(l1_waybel_0(esk2_0,esk1_0)&(m1_yellow_6(esk3_0,esk1_0,esk2_0)&(m1_subset_1(esk4_0,u1_struct_0(esk2_0))&(m1_subset_1(esk5_0,u1_struct_0(esk2_0))&(m1_subset_1(esk6_0,u1_struct_0(esk3_0))&(m1_subset_1(esk7_0,u1_struct_0(esk3_0))&((esk4_0=esk6_0&(esk5_0=esk7_0&r1_orders_2(esk3_0,esk6_0,esk7_0)))&~r1_orders_2(esk2_0,esk4_0,esk5_0))))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
fof(c_0_20, plain, ![X21, X22]:(~l1_struct_0(X21)|(~l1_waybel_0(X22,X21)|l1_orders_2(X22))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_waybel_0])])])).
cnf(c_0_21, plain, (r1_tarski(u1_orders_2(X1),u1_orders_2(X2))|~m1_yellow_0(X1,X2)|~l1_orders_2(X1)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (l1_orders_2(X2)|~l1_orders_2(X1)|~m1_yellow_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (m1_yellow_0(X1,X2)|~m1_yellow_6(X1,X3,X2)|~l1_waybel_0(X2,X3)|~l1_struct_0(X3)), inference(csr,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_24, negated_conjecture, (m1_yellow_6(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, negated_conjecture, (l1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (l1_orders_2(X2)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_28, plain, ![X19, X20]:((~m1_subset_1(X19,k1_zfmisc_1(X20))|r1_tarski(X19,X20))&(~r1_tarski(X19,X20)|m1_subset_1(X19,k1_zfmisc_1(X20)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_29, plain, (r1_tarski(u1_orders_2(X1),u1_orders_2(X2))|~m1_yellow_0(X1,X2)|~l1_orders_2(X2)), inference(csr,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_30, negated_conjecture, (m1_yellow_0(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25]), c_0_26])])).
cnf(c_0_31, negated_conjecture, (l1_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_25]), c_0_26])])).
fof(c_0_32, plain, ![X45, X46]:(~m1_subset_1(X45,X46)|(v1_xboole_0(X46)|r2_hidden(X45,X46))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_33, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34, negated_conjecture, (r1_tarski(u1_orders_2(esk3_0),u1_orders_2(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])])).
fof(c_0_35, plain, ![X18]:~v1_xboole_0(k1_zfmisc_1(X18)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
fof(c_0_36, plain, ![X15, X16, X17]:((~r1_orders_2(X15,X16,X17)|r2_hidden(k4_tarski(X16,X17),u1_orders_2(X15))|~m1_subset_1(X17,u1_struct_0(X15))|~m1_subset_1(X16,u1_struct_0(X15))|~l1_orders_2(X15))&(~r2_hidden(k4_tarski(X16,X17),u1_orders_2(X15))|r1_orders_2(X15,X16,X17)|~m1_subset_1(X17,u1_struct_0(X15))|~m1_subset_1(X16,u1_struct_0(X15))|~l1_orders_2(X15))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_orders_2])])])])).
cnf(c_0_37, negated_conjecture, (l1_waybel_0(esk3_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_24]), c_0_25]), c_0_26])])).
fof(c_0_38, plain, ![X34, X35, X36, X38, X39, X40, X41, X43]:((((r2_hidden(X36,esk8_3(X34,X35,X36))|~r2_hidden(X36,X35)|X35!=k3_tarski(X34))&(r2_hidden(esk8_3(X34,X35,X36),X34)|~r2_hidden(X36,X35)|X35!=k3_tarski(X34)))&(~r2_hidden(X38,X39)|~r2_hidden(X39,X34)|r2_hidden(X38,X35)|X35!=k3_tarski(X34)))&((~r2_hidden(esk9_2(X40,X41),X41)|(~r2_hidden(esk9_2(X40,X41),X43)|~r2_hidden(X43,X40))|X41=k3_tarski(X40))&((r2_hidden(esk9_2(X40,X41),esk10_2(X40,X41))|r2_hidden(esk9_2(X40,X41),X41)|X41=k3_tarski(X40))&(r2_hidden(esk10_2(X40,X41),X40)|r2_hidden(esk9_2(X40,X41),X41)|X41=k3_tarski(X40))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_tarski])])])])])])).
cnf(c_0_39, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40, negated_conjecture, (m1_subset_1(u1_orders_2(esk3_0),k1_zfmisc_1(u1_orders_2(esk2_0)))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_41, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_42, plain, ![X30]:k3_tarski(k1_zfmisc_1(X30))=X30, inference(variable_rename,[status(thm)],[t99_zfmisc_1])).
cnf(c_0_43, plain, (r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44, negated_conjecture, (m1_subset_1(esk7_0,u1_struct_0(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_45, negated_conjecture, (l1_orders_2(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_37]), c_0_26])])).
cnf(c_0_46, negated_conjecture, (r1_orders_2(esk3_0,esk6_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_47, negated_conjecture, (esk4_0=esk6_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_48, negated_conjecture, (m1_subset_1(esk6_0,u1_struct_0(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_49, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X2,X3)|X4!=k3_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, negated_conjecture, (r2_hidden(u1_orders_2(esk3_0),k1_zfmisc_1(u1_orders_2(esk2_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])).
cnf(c_0_51, plain, (k3_tarski(k1_zfmisc_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_52, negated_conjecture, (r2_hidden(k4_tarski(X1,esk7_0),u1_orders_2(esk3_0))|~r1_orders_2(esk3_0,X1,esk7_0)|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])])).
cnf(c_0_53, negated_conjecture, (r1_orders_2(esk3_0,esk4_0,esk7_0)), inference(rw,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk3_0))), inference(rw,[status(thm)],[c_0_48, c_0_47])).
cnf(c_0_55, negated_conjecture, (r2_hidden(X1,X2)|X2!=u1_orders_2(esk2_0)|~r2_hidden(X1,u1_orders_2(esk3_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51])).
cnf(c_0_56, negated_conjecture, (r2_hidden(k4_tarski(esk4_0,esk7_0),u1_orders_2(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54])])).
cnf(c_0_57, negated_conjecture, (r2_hidden(k4_tarski(esk4_0,esk7_0),X1)|X1!=u1_orders_2(esk2_0)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_58, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_59, negated_conjecture, (esk5_0=esk7_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_60, negated_conjecture, (~r1_orders_2(esk2_0,esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_61, plain, (r1_orders_2(X3,X1,X2)|~r2_hidden(k4_tarski(X1,X2),u1_orders_2(X3))|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_62, negated_conjecture, (r2_hidden(k4_tarski(esk4_0,esk7_0),u1_orders_2(esk2_0))), inference(er,[status(thm)],[c_0_57])).
cnf(c_0_63, negated_conjecture, (m1_subset_1(esk7_0,u1_struct_0(esk2_0))), inference(rw,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_64, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_65, negated_conjecture, (~r1_orders_2(esk2_0,esk4_0,esk7_0)), inference(rw,[status(thm)],[c_0_60, c_0_59])).
cnf(c_0_66, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_31]), c_0_63]), c_0_64])]), c_0_65]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 67
# Proof object clause steps            : 42
# Proof object formula steps           : 25
# Proof object conjectures             : 31
# Proof object clause conjectures      : 28
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 29
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 33
# Processed clauses                    : 95
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 95
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 67
# ...of the previous two non-trivial   : 58
# Contextual simplify-reflections      : 4
# Paramodulations                      : 62
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 62
#    Positive orientable unit clauses  : 24
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 36
# Current number of unprocessed clauses: 26
# ...number of literals in the above   : 93
# Current number of archived formulas  : 0
# Current number of archived clauses   : 33
# Clause-clause subsumption calls (NU) : 232
# Rec. Clause-clause subsumption calls : 70
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 8
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3487

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28644 pages
