# No SInE strategy applied
# Auto-Mode selected heuristic H_____047_B31_F1_PI_AE_R4_CS_SP_S4c
# and selection function SelectCQPrecWNTNp.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/779ade09cc94481f98481026cdb54d3b.E_input', redefinition_k6_domain_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/779ade09cc94481f98481026cdb54d3b.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/779ade09cc94481f98481026cdb54d3b.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/779ade09cc94481f98481026cdb54d3b.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/779ade09cc94481f98481026cdb54d3b.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/779ade09cc94481f98481026cdb54d3b.E_input', t73_enumset1)).
fof(t50_orders_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>~(r2_hidden(X2,k2_orders_2(X1,k6_domain_1(u1_struct_0(X1),X2)))))), file('data/MPTP/proofs/779ade09cc94481f98481026cdb54d3b.E_input', t50_orders_2)).
fof(d4_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:(X7=k4_enumset1(X1,X2,X3,X4,X5,X6)<=>![X8]:(r2_hidden(X8,X7)<=>~((~(X8=X1)&(~(X8=X2)&(~(X8=X3)&(~(X8=X4)&(~(X8=X5)&~(X8=X6))))))))), file('data/MPTP/proofs/779ade09cc94481f98481026cdb54d3b.E_input', d4_enumset1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/779ade09cc94481f98481026cdb54d3b.E_input', fc2_struct_0)).
fof(t49_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>~((r2_hidden(X2,X3)&r2_hidden(X2,k2_orders_2(X1,X3))))))), file('data/MPTP/proofs/779ade09cc94481f98481026cdb54d3b.E_input', t49_orders_2)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/779ade09cc94481f98481026cdb54d3b.E_input', t4_subset)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/779ade09cc94481f98481026cdb54d3b.E_input', dt_l1_orders_2)).
fof(t35_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v6_orders_2(k6_domain_1(u1_struct_0(X1),X2),X1)&m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/779ade09cc94481f98481026cdb54d3b.E_input', t35_orders_2)).
fof(c_0_13, plain, ![X9, X10]:(v1_xboole_0(X9)|~m1_subset_1(X10,X9)|k6_domain_1(X9,X10)=k1_tarski(X10)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_14, plain, ![X11]:k2_tarski(X11,X11)=k1_tarski(X11), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_15, plain, ![X12, X13]:k1_enumset1(X12,X12,X13)=k2_tarski(X12,X13), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_16, plain, ![X14, X15, X16]:k2_enumset1(X14,X14,X15,X16)=k1_enumset1(X14,X15,X16), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_17, plain, ![X17, X18, X19, X20]:k3_enumset1(X17,X17,X18,X19,X20)=k2_enumset1(X17,X18,X19,X20), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_18, plain, ![X21, X22, X23, X24, X25]:k4_enumset1(X21,X21,X22,X23,X24,X25)=k3_enumset1(X21,X22,X23,X24,X25), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_19, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>~(r2_hidden(X2,k2_orders_2(X1,k6_domain_1(u1_struct_0(X1),X2))))))), inference(assume_negation,[status(cth)],[t50_orders_2])).
cnf(c_0_20, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_26, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&r2_hidden(esk2_0,k2_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_19])])])])).
fof(c_0_27, plain, ![X28, X29, X30, X31, X32, X33, X34, X35, X36, X37, X38, X39, X40, X41, X42, X43]:(((~r2_hidden(X35,X34)|(X35=X28|(X35=X29|(X35=X30|(X35=X31|(X35=X32|X35=X33)))))|X34!=k4_enumset1(X28,X29,X30,X31,X32,X33))&((X36!=X28|r2_hidden(X36,X34)|X34!=k4_enumset1(X28,X29,X30,X31,X32,X33))&((X36!=X29|r2_hidden(X36,X34)|X34!=k4_enumset1(X28,X29,X30,X31,X32,X33))&((X36!=X30|r2_hidden(X36,X34)|X34!=k4_enumset1(X28,X29,X30,X31,X32,X33))&((X36!=X31|r2_hidden(X36,X34)|X34!=k4_enumset1(X28,X29,X30,X31,X32,X33))&((X36!=X32|r2_hidden(X36,X34)|X34!=k4_enumset1(X28,X29,X30,X31,X32,X33))&(X36!=X33|r2_hidden(X36,X34)|X34!=k4_enumset1(X28,X29,X30,X31,X32,X33))))))))&(((esk3_7(X37,X38,X39,X40,X41,X42,X43)!=X37|~r2_hidden(esk3_7(X37,X38,X39,X40,X41,X42,X43),X43)|X43=k4_enumset1(X37,X38,X39,X40,X41,X42))&((esk3_7(X37,X38,X39,X40,X41,X42,X43)!=X38|~r2_hidden(esk3_7(X37,X38,X39,X40,X41,X42,X43),X43)|X43=k4_enumset1(X37,X38,X39,X40,X41,X42))&((esk3_7(X37,X38,X39,X40,X41,X42,X43)!=X39|~r2_hidden(esk3_7(X37,X38,X39,X40,X41,X42,X43),X43)|X43=k4_enumset1(X37,X38,X39,X40,X41,X42))&((esk3_7(X37,X38,X39,X40,X41,X42,X43)!=X40|~r2_hidden(esk3_7(X37,X38,X39,X40,X41,X42,X43),X43)|X43=k4_enumset1(X37,X38,X39,X40,X41,X42))&((esk3_7(X37,X38,X39,X40,X41,X42,X43)!=X41|~r2_hidden(esk3_7(X37,X38,X39,X40,X41,X42,X43),X43)|X43=k4_enumset1(X37,X38,X39,X40,X41,X42))&(esk3_7(X37,X38,X39,X40,X41,X42,X43)!=X42|~r2_hidden(esk3_7(X37,X38,X39,X40,X41,X42,X43),X43)|X43=k4_enumset1(X37,X38,X39,X40,X41,X42)))))))&(r2_hidden(esk3_7(X37,X38,X39,X40,X41,X42,X43),X43)|(esk3_7(X37,X38,X39,X40,X41,X42,X43)=X37|(esk3_7(X37,X38,X39,X40,X41,X42,X43)=X38|(esk3_7(X37,X38,X39,X40,X41,X42,X43)=X39|(esk3_7(X37,X38,X39,X40,X41,X42,X43)=X40|(esk3_7(X37,X38,X39,X40,X41,X42,X43)=X41|esk3_7(X37,X38,X39,X40,X41,X42,X43)=X42)))))|X43=k4_enumset1(X37,X38,X39,X40,X41,X42)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_enumset1])])])])])])])).
fof(c_0_28, plain, ![X45]:(v2_struct_0(X45)|~l1_struct_0(X45)|~v1_xboole_0(u1_struct_0(X45))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_29, plain, (k6_domain_1(X1,X2)=k4_enumset1(X2,X2,X2,X2,X2,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25])).
cnf(c_0_30, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_31, plain, ![X46, X47, X48]:(v2_struct_0(X46)|(~v3_orders_2(X46)|(~v4_orders_2(X46)|(~v5_orders_2(X46)|~l1_orders_2(X46))))|(~m1_subset_1(X47,u1_struct_0(X46))|(~m1_subset_1(X48,k1_zfmisc_1(u1_struct_0(X46)))|(~r2_hidden(X47,X48)|~r2_hidden(X47,k2_orders_2(X46,X48)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t49_orders_2])])])])).
fof(c_0_32, plain, ![X49, X50, X51]:(~r2_hidden(X49,X50)|~m1_subset_1(X50,k1_zfmisc_1(X51))|m1_subset_1(X49,X51)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_33, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k4_enumset1(X4,X5,X6,X7,X8,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_34, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k4_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_36, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_37, plain, ![X52]:(~l1_orders_2(X52)|l1_struct_0(X52)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
fof(c_0_38, plain, ![X53, X54]:((v6_orders_2(k6_domain_1(u1_struct_0(X53),X54),X53)|~m1_subset_1(X54,u1_struct_0(X53))|(v2_struct_0(X53)|(~v3_orders_2(X53)|~l1_orders_2(X53))))&(m1_subset_1(k6_domain_1(u1_struct_0(X53),X54),k1_zfmisc_1(u1_struct_0(X53)))|~m1_subset_1(X54,u1_struct_0(X53))|(v2_struct_0(X53)|(~v3_orders_2(X53)|~l1_orders_2(X53))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t35_orders_2])])])])])).
cnf(c_0_39, plain, (v2_struct_0(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,X3)|~r2_hidden(X2,k2_orders_2(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_40, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_41, plain, (r2_hidden(X1,X2)|X2!=k4_enumset1(X3,X4,X5,X6,X7,X1)), inference(er,[status(thm)],[c_0_33])).
cnf(c_0_42, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k4_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])).
cnf(c_0_43, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_44, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_45, plain, (m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_46, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_47, plain, (v2_struct_0(X1)|~r2_hidden(X2,k2_orders_2(X1,X3))|~r2_hidden(X2,X3)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(csr,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_48, plain, (r2_hidden(X1,k4_enumset1(X2,X3,X4,X5,X6,X1))), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_49, negated_conjecture, (r2_hidden(esk2_0,k2_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_50, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k4_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])])).
cnf(c_0_51, negated_conjecture, (m1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_30]), c_0_44]), c_0_46])]), c_0_36])).
cnf(c_0_52, plain, (v2_struct_0(X1)|~r2_hidden(X2,k2_orders_2(X1,k4_enumset1(X3,X4,X5,X6,X7,X2)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k4_enumset1(X3,X4,X5,X6,X7,X2),k1_zfmisc_1(u1_struct_0(X1)))), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_53, negated_conjecture, (r2_hidden(esk2_0,k2_orders_2(esk1_0,k4_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))), inference(rw,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_54, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_55, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_56, negated_conjecture, (m1_subset_1(k4_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(rw,[status(thm)],[c_0_51, c_0_50])).
cnf(c_0_57, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_44]), c_0_54]), c_0_55]), c_0_46]), c_0_56])]), c_0_36]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 58
# Proof object clause steps            : 31
# Proof object formula steps           : 27
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 23
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 28
# Processed clauses                    : 66
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 65
# Other redundant clauses eliminated   : 13
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 5
# Generated clauses                    : 186
# ...of the previous two non-trivial   : 173
# Contextual simplify-reflections      : 1
# Paramodulations                      : 166
# Factorizations                       : 0
# Equation resolutions                 : 20
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 54
#    Positive orientable unit clauses  : 15
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 38
# Current number of unprocessed clauses: 128
# ...number of literals in the above   : 1123
# Current number of archived formulas  : 0
# Current number of archived clauses   : 10
# Clause-clause subsumption calls (NU) : 241
# Rec. Clause-clause subsumption calls : 149
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 6
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 16
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5884

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.000 s
# Total time               : 0.013 s
# Maximum resident set size: 28700 pages
