# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(r1_tarski(X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X3,X4)=>r1_lattice3(X1,X2,X4))&(r2_lattice3(X1,X3,X4)=>r2_lattice3(X1,X2,X4)))))), file('data/MPTP/proofs/779d6acf173d4b8f8c9134b68c086854.E_input', t9_yellow_0)).
fof(t32_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v3_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(X2=k1_yellow_0(X1,X3)<=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))))), file('data/MPTP/proofs/779d6acf173d4b8f8c9134b68c086854.E_input', t32_yellow_0)).
fof(t3_waybel_7, conjecture, ![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v3_lattice3(X1)&l1_orders_2(X1)))))))=>![X2, X3]:(r1_tarski(X2,X3)=>(r3_orders_2(X1,k1_yellow_0(X1,X2),k1_yellow_0(X1,X3))&r1_orders_2(X1,k2_yellow_0(X1,X3),k2_yellow_0(X1,X2))))), file('data/MPTP/proofs/779d6acf173d4b8f8c9134b68c086854.E_input', t3_waybel_7)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/779d6acf173d4b8f8c9134b68c086854.E_input', dt_k1_yellow_0)).
fof(d10_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_yellow_0(X1,X2)=>(X3=k2_yellow_0(X1,X2)<=>(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))))))), file('data/MPTP/proofs/779d6acf173d4b8f8c9134b68c086854.E_input', d10_yellow_0)).
fof(dt_k2_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/779d6acf173d4b8f8c9134b68c086854.E_input', dt_k2_yellow_0)).
fof(cc1_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v1_lattice3(X1)=>~(v2_struct_0(X1)))), file('data/MPTP/proofs/779d6acf173d4b8f8c9134b68c086854.E_input', cc1_lattice3)).
fof(t17_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v3_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(r1_yellow_0(X1,X2)&r2_yellow_0(X1,X2))), file('data/MPTP/proofs/779d6acf173d4b8f8c9134b68c086854.E_input', t17_yellow_0)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/779d6acf173d4b8f8c9134b68c086854.E_input', redefinition_r3_orders_2)).
fof(t33_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v3_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(X2=k2_yellow_0(X1,X3)<=>(r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2))))))), file('data/MPTP/proofs/779d6acf173d4b8f8c9134b68c086854.E_input', t33_yellow_0)).
fof(c_0_10, plain, ![X35, X36, X37, X38]:((~r1_lattice3(X35,X37,X38)|r1_lattice3(X35,X36,X38)|~m1_subset_1(X38,u1_struct_0(X35))|~r1_tarski(X36,X37)|~l1_orders_2(X35))&(~r2_lattice3(X35,X37,X38)|r2_lattice3(X35,X36,X38)|~m1_subset_1(X38,u1_struct_0(X35))|~r1_tarski(X36,X37)|~l1_orders_2(X35))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_yellow_0])])])])).
fof(c_0_11, plain, ![X17, X18, X19, X20, X21]:(((r2_lattice3(X17,X19,X18)|X18!=k1_yellow_0(X17,X19)|~m1_subset_1(X18,u1_struct_0(X17))|(v2_struct_0(X17)|(~v5_orders_2(X17)|(~v3_lattice3(X17)|~l1_orders_2(X17)))))&(~m1_subset_1(X20,u1_struct_0(X17))|(~r2_lattice3(X17,X19,X20)|r1_orders_2(X17,X18,X20))|X18!=k1_yellow_0(X17,X19)|~m1_subset_1(X18,u1_struct_0(X17))|(v2_struct_0(X17)|(~v5_orders_2(X17)|(~v3_lattice3(X17)|~l1_orders_2(X17))))))&((m1_subset_1(esk4_3(X17,X18,X21),u1_struct_0(X17))|~r2_lattice3(X17,X21,X18)|X18=k1_yellow_0(X17,X21)|~m1_subset_1(X18,u1_struct_0(X17))|(v2_struct_0(X17)|(~v5_orders_2(X17)|(~v3_lattice3(X17)|~l1_orders_2(X17)))))&((r2_lattice3(X17,X21,esk4_3(X17,X18,X21))|~r2_lattice3(X17,X21,X18)|X18=k1_yellow_0(X17,X21)|~m1_subset_1(X18,u1_struct_0(X17))|(v2_struct_0(X17)|(~v5_orders_2(X17)|(~v3_lattice3(X17)|~l1_orders_2(X17)))))&(~r1_orders_2(X17,X18,esk4_3(X17,X18,X21))|~r2_lattice3(X17,X21,X18)|X18=k1_yellow_0(X17,X21)|~m1_subset_1(X18,u1_struct_0(X17))|(v2_struct_0(X17)|(~v5_orders_2(X17)|(~v3_lattice3(X17)|~l1_orders_2(X17)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t32_yellow_0])])])])])])])).
fof(c_0_12, negated_conjecture, ~(![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v3_lattice3(X1)&l1_orders_2(X1)))))))=>![X2, X3]:(r1_tarski(X2,X3)=>(r3_orders_2(X1,k1_yellow_0(X1,X2),k1_yellow_0(X1,X3))&r1_orders_2(X1,k2_yellow_0(X1,X3),k2_yellow_0(X1,X2)))))), inference(assume_negation,[status(cth)],[t3_waybel_7])).
cnf(c_0_13, plain, (r2_lattice3(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_tarski(X4,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (r2_lattice3(X1,X2,X3)|v2_struct_0(X1)|X3!=k1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~v3_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_15, plain, ![X12, X13]:(~l1_orders_2(X12)|m1_subset_1(k1_yellow_0(X12,X13),u1_struct_0(X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
fof(c_0_16, plain, ![X23, X24, X25, X26]:(((r1_lattice3(X23,X24,X25)|X25!=k2_yellow_0(X23,X24)|~r2_yellow_0(X23,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))&(~m1_subset_1(X26,u1_struct_0(X23))|(~r1_lattice3(X23,X24,X26)|r1_orders_2(X23,X26,X25))|X25!=k2_yellow_0(X23,X24)|~r2_yellow_0(X23,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23)))&((m1_subset_1(esk5_3(X23,X24,X25),u1_struct_0(X23))|~r1_lattice3(X23,X24,X25)|X25=k2_yellow_0(X23,X24)|~r2_yellow_0(X23,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))&((r1_lattice3(X23,X24,esk5_3(X23,X24,X25))|~r1_lattice3(X23,X24,X25)|X25=k2_yellow_0(X23,X24)|~r2_yellow_0(X23,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))&(~r1_orders_2(X23,esk5_3(X23,X24,X25),X25)|~r1_lattice3(X23,X24,X25)|X25=k2_yellow_0(X23,X24)|~r2_yellow_0(X23,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_yellow_0])])])])])).
fof(c_0_17, plain, ![X8, X9]:(~l1_orders_2(X8)|m1_subset_1(k2_yellow_0(X8,X9),u1_struct_0(X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_yellow_0])])).
fof(c_0_18, plain, ![X34]:(~l1_orders_2(X34)|(~v1_lattice3(X34)|~v2_struct_0(X34))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattice3])])])).
fof(c_0_19, negated_conjecture, ((v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_lattice3(esk1_0)&(v2_lattice3(esk1_0)&(v3_lattice3(esk1_0)&l1_orders_2(esk1_0)))))))&(r1_tarski(esk2_0,esk3_0)&(~r3_orders_2(esk1_0,k1_yellow_0(esk1_0,esk2_0),k1_yellow_0(esk1_0,esk3_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),k2_yellow_0(esk1_0,esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
cnf(c_0_20, plain, (r2_lattice3(X1,X2,X3)|v2_struct_0(X1)|X3!=k1_yellow_0(X1,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_tarski(X2,X4)|~l1_orders_2(X1)|~v3_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_21, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (r1_lattice3(X1,X2,X3)|X3!=k2_yellow_0(X1,X2)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_24, plain, ![X10, X11]:((r1_yellow_0(X10,X11)|(v2_struct_0(X10)|(~v5_orders_2(X10)|(~v3_lattice3(X10)|~l1_orders_2(X10)))))&(r2_yellow_0(X10,X11)|(v2_struct_0(X10)|(~v5_orders_2(X10)|(~v3_lattice3(X10)|~l1_orders_2(X10)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_yellow_0])])])])])).
fof(c_0_25, plain, ![X14, X15, X16]:((~r3_orders_2(X14,X15,X16)|r1_orders_2(X14,X15,X16)|(v2_struct_0(X14)|(~v3_orders_2(X14)|~l1_orders_2(X14))|(~m1_subset_1(X15,u1_struct_0(X14))|~m1_subset_1(X16,u1_struct_0(X14)))))&(~r1_orders_2(X14,X15,X16)|r3_orders_2(X14,X15,X16)|(v2_struct_0(X14)|(~v3_orders_2(X14)|~l1_orders_2(X14))|(~m1_subset_1(X15,u1_struct_0(X14))|~m1_subset_1(X16,u1_struct_0(X14)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_26, plain, (~l1_orders_2(X1)|~v1_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, negated_conjecture, (v1_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (r2_lattice3(X1,X2,k1_yellow_0(X1,X3))|v2_struct_0(X1)|~r1_tarski(X2,X3)|~l1_orders_2(X1)|~v3_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_20]), c_0_21])).
cnf(c_0_30, negated_conjecture, (r1_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_31, plain, (r1_lattice3(X1,X2,k2_yellow_0(X1,X2))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_22]), c_0_23])).
cnf(c_0_32, plain, (r2_yellow_0(X1,X2)|v2_struct_0(X1)|~v5_orders_2(X1)|~v3_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, negated_conjecture, (~r3_orders_2(esk1_0,k1_yellow_0(esk1_0,esk2_0),k1_yellow_0(esk1_0,esk3_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),k2_yellow_0(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_34, plain, (r3_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_36, negated_conjecture, (~v2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])])).
cnf(c_0_37, plain, (r1_orders_2(X2,X4,X1)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~v5_orders_2(X2)|~v3_lattice3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_38, negated_conjecture, (r2_lattice3(X1,esk2_0,k1_yellow_0(X1,esk3_0))|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_39, plain, (r1_lattice3(X1,X4,X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_tarski(X4,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_40, plain, (r1_lattice3(X1,X2,k2_yellow_0(X1,X2))|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_41, negated_conjecture, (~m1_subset_1(k1_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),k2_yellow_0(esk1_0,esk2_0))|~r1_orders_2(esk1_0,k1_yellow_0(esk1_0,esk2_0),k1_yellow_0(esk1_0,esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_28]), c_0_35])]), c_0_36])).
cnf(c_0_42, negated_conjecture, (v2_struct_0(X1)|r1_orders_2(X1,X2,k1_yellow_0(X1,esk3_0))|X2!=k1_yellow_0(X1,esk2_0)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_21])).
cnf(c_0_43, negated_conjecture, (v3_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_44, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_45, plain, ![X28, X29, X30, X31, X32]:(((r1_lattice3(X28,X30,X29)|X29!=k2_yellow_0(X28,X30)|~m1_subset_1(X29,u1_struct_0(X28))|(v2_struct_0(X28)|(~v5_orders_2(X28)|(~v3_lattice3(X28)|~l1_orders_2(X28)))))&(~m1_subset_1(X31,u1_struct_0(X28))|(~r1_lattice3(X28,X30,X31)|r1_orders_2(X28,X31,X29))|X29!=k2_yellow_0(X28,X30)|~m1_subset_1(X29,u1_struct_0(X28))|(v2_struct_0(X28)|(~v5_orders_2(X28)|(~v3_lattice3(X28)|~l1_orders_2(X28))))))&((m1_subset_1(esk6_3(X28,X29,X32),u1_struct_0(X28))|~r1_lattice3(X28,X32,X29)|X29=k2_yellow_0(X28,X32)|~m1_subset_1(X29,u1_struct_0(X28))|(v2_struct_0(X28)|(~v5_orders_2(X28)|(~v3_lattice3(X28)|~l1_orders_2(X28)))))&((r1_lattice3(X28,X32,esk6_3(X28,X29,X32))|~r1_lattice3(X28,X32,X29)|X29=k2_yellow_0(X28,X32)|~m1_subset_1(X29,u1_struct_0(X28))|(v2_struct_0(X28)|(~v5_orders_2(X28)|(~v3_lattice3(X28)|~l1_orders_2(X28)))))&(~r1_orders_2(X28,esk6_3(X28,X29,X32),X29)|~r1_lattice3(X28,X32,X29)|X29=k2_yellow_0(X28,X32)|~m1_subset_1(X29,u1_struct_0(X28))|(v2_struct_0(X28)|(~v5_orders_2(X28)|(~v3_lattice3(X28)|~l1_orders_2(X28)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t33_yellow_0])])])])])])])).
cnf(c_0_46, plain, (r1_lattice3(X1,X2,k2_yellow_0(X1,X3))|v2_struct_0(X1)|~r1_tarski(X2,X3)|~l1_orders_2(X1)|~v3_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_23])).
cnf(c_0_47, negated_conjecture, (~m1_subset_1(k1_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),k2_yellow_0(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_28]), c_0_43]), c_0_44])]), c_0_36])).
cnf(c_0_48, plain, (r1_orders_2(X2,X1,X4)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|X4!=k2_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~v5_orders_2(X2)|~v3_lattice3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_49, negated_conjecture, (r1_lattice3(X1,esk2_0,k2_yellow_0(X1,esk3_0))|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_46, c_0_30])).
cnf(c_0_50, negated_conjecture, (~m1_subset_1(k1_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),k2_yellow_0(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_21]), c_0_28])])).
cnf(c_0_51, negated_conjecture, (v2_struct_0(X1)|r1_orders_2(X1,k2_yellow_0(X1,esk3_0),X2)|X2!=k2_yellow_0(X1,esk2_0)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_23])).
cnf(c_0_52, negated_conjecture, (~m1_subset_1(k1_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_28]), c_0_43]), c_0_44])]), c_0_36])).
cnf(c_0_53, negated_conjecture, (~m1_subset_1(k2_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_21]), c_0_28])])).
cnf(c_0_54, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_23]), c_0_28])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 55
# Proof object clause steps            : 34
# Proof object formula steps           : 21
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 27
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 33
# Processed clauses                    : 86
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 84
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 0
# Generated clauses                    : 36
# ...of the previous two non-trivial   : 34
# Contextual simplify-reflections      : 7
# Paramodulations                      : 32
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 49
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 39
# Current number of unprocessed clauses: 14
# ...number of literals in the above   : 136
# Current number of archived formulas  : 0
# Current number of archived clauses   : 35
# Clause-clause subsumption calls (NU) : 1249
# Rec. Clause-clause subsumption calls : 32
# Non-unit clause-clause subsumptions  : 11
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4508

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.003 s
# Total time               : 0.012 s
# Maximum resident set size: 28580 pages
