# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/77ca838c5c514663802b886cf28e662e.E_input', d4_xboole_0)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/77ca838c5c514663802b886cf28e662e.E_input', d5_xboole_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/77ca838c5c514663802b886cf28e662e.E_input', d1_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/77ca838c5c514663802b886cf28e662e.E_input', commutativity_k3_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/77ca838c5c514663802b886cf28e662e.E_input', d3_tarski)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/77ca838c5c514663802b886cf28e662e.E_input', t7_xboole_0)).
fof(t134_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4)))), file('data/MPTP/proofs/77ca838c5c514663802b886cf28e662e.E_input', t134_zfmisc_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/77ca838c5c514663802b886cf28e662e.E_input', t113_zfmisc_1)).
fof(t122_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k3_xboole_0(X1,X2),X3)=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&k2_zfmisc_1(X3,k3_xboole_0(X1,X2))=k3_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))), file('data/MPTP/proofs/77ca838c5c514663802b886cf28e662e.E_input', t122_zfmisc_1)).
fof(t123_zfmisc_1, axiom, ![X1, X2, X3, X4]:k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)), file('data/MPTP/proofs/77ca838c5c514663802b886cf28e662e.E_input', t123_zfmisc_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/77ca838c5c514663802b886cf28e662e.E_input', t28_xboole_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/77ca838c5c514663802b886cf28e662e.E_input', idempotence_k3_xboole_0)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/77ca838c5c514663802b886cf28e662e.E_input', d10_xboole_0)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/77ca838c5c514663802b886cf28e662e.E_input', l32_xboole_1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/77ca838c5c514663802b886cf28e662e.E_input', t4_xboole_0)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/77ca838c5c514663802b886cf28e662e.E_input', t65_xboole_1)).
fof(c_0_16, plain, ![X5, X6, X7, X8, X9, X10, X11, X12]:((((r2_hidden(X8,X5)|~r2_hidden(X8,X7)|X7!=k3_xboole_0(X5,X6))&(r2_hidden(X8,X6)|~r2_hidden(X8,X7)|X7!=k3_xboole_0(X5,X6)))&(~r2_hidden(X9,X5)|~r2_hidden(X9,X6)|r2_hidden(X9,X7)|X7!=k3_xboole_0(X5,X6)))&((~r2_hidden(esk1_3(X10,X11,X12),X12)|(~r2_hidden(esk1_3(X10,X11,X12),X10)|~r2_hidden(esk1_3(X10,X11,X12),X11))|X12=k3_xboole_0(X10,X11))&((r2_hidden(esk1_3(X10,X11,X12),X10)|r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k3_xboole_0(X10,X11))&(r2_hidden(esk1_3(X10,X11,X12),X11)|r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k3_xboole_0(X10,X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_17, plain, ![X14, X15, X16, X17, X18, X19, X20, X21]:((((r2_hidden(X17,X14)|~r2_hidden(X17,X16)|X16!=k4_xboole_0(X14,X15))&(~r2_hidden(X17,X15)|~r2_hidden(X17,X16)|X16!=k4_xboole_0(X14,X15)))&(~r2_hidden(X18,X14)|r2_hidden(X18,X15)|r2_hidden(X18,X16)|X16!=k4_xboole_0(X14,X15)))&((~r2_hidden(esk2_3(X19,X20,X21),X21)|(~r2_hidden(esk2_3(X19,X20,X21),X19)|r2_hidden(esk2_3(X19,X20,X21),X20))|X21=k4_xboole_0(X19,X20))&((r2_hidden(esk2_3(X19,X20,X21),X19)|r2_hidden(esk2_3(X19,X20,X21),X21)|X21=k4_xboole_0(X19,X20))&(~r2_hidden(esk2_3(X19,X20,X21),X20)|r2_hidden(esk2_3(X19,X20,X21),X21)|X21=k4_xboole_0(X19,X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_18, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_19, plain, ![X23, X24, X25]:((~v1_xboole_0(X23)|~r2_hidden(X24,X23))&(r2_hidden(esk3_1(X25),X25)|v1_xboole_0(X25))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_20, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_18])).
cnf(c_0_22, plain, (r2_hidden(esk3_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_23, plain, ![X27, X28]:k3_xboole_0(X27,X28)=k3_xboole_0(X28,X27), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_24, plain, ![X29, X30, X31, X32, X33]:((~r1_tarski(X29,X30)|(~r2_hidden(X31,X29)|r2_hidden(X31,X30)))&((r2_hidden(esk4_2(X32,X33),X32)|r1_tarski(X32,X33))&(~r2_hidden(esk4_2(X32,X33),X33)|r1_tarski(X32,X33)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_25, plain, ![X35]:(X35=k1_xboole_0|r2_hidden(esk5_1(X35),X35)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_26, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|r2_hidden(esk3_1(k3_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_30, negated_conjecture, ~(![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4))))), inference(assume_negation,[status(cth)],[t134_zfmisc_1])).
cnf(c_0_31, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_32, plain, (X1=k1_xboole_0|r2_hidden(esk5_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (v1_xboole_0(k3_xboole_0(X1,k4_xboole_0(X2,X3)))|~r2_hidden(esk3_1(k3_xboole_0(X1,k4_xboole_0(X2,X3))),X3)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_34, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|r2_hidden(esk3_1(k3_xboole_0(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
fof(c_0_35, plain, ![X41, X42]:((k2_zfmisc_1(X41,X42)!=k1_xboole_0|(X41=k1_xboole_0|X42=k1_xboole_0))&((X41!=k1_xboole_0|k2_zfmisc_1(X41,X42)=k1_xboole_0)&(X42!=k1_xboole_0|k2_zfmisc_1(X41,X42)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_36, plain, (r1_tarski(X1,X2)|~r2_hidden(esk4_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_37, plain, (r1_tarski(k3_xboole_0(X1,X2),X3)|r2_hidden(esk4_2(k3_xboole_0(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_21, c_0_29])).
fof(c_0_38, plain, ![X43, X44, X45]:(k2_zfmisc_1(k3_xboole_0(X43,X44),X45)=k3_xboole_0(k2_zfmisc_1(X43,X45),k2_zfmisc_1(X44,X45))&k2_zfmisc_1(X45,k3_xboole_0(X43,X44))=k3_xboole_0(k2_zfmisc_1(X45,X43),k2_zfmisc_1(X45,X44))), inference(variable_rename,[status(thm)],[t122_zfmisc_1])).
fof(c_0_39, negated_conjecture, (k2_zfmisc_1(esk6_0,esk7_0)=k2_zfmisc_1(esk8_0,esk9_0)&(esk6_0!=k1_xboole_0&(esk7_0!=k1_xboole_0&(esk6_0!=esk8_0|esk7_0!=esk9_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])).
cnf(c_0_40, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_41, plain, ![X46, X47, X48, X49]:k2_zfmisc_1(k3_xboole_0(X46,X47),k3_xboole_0(X48,X49))=k3_xboole_0(k2_zfmisc_1(X46,X48),k2_zfmisc_1(X47,X49)), inference(variable_rename,[status(thm)],[t123_zfmisc_1])).
cnf(c_0_42, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_43, plain, (v1_xboole_0(k3_xboole_0(X1,k4_xboole_0(X2,X1)))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_44, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_45, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_46, plain, (k2_zfmisc_1(X1,k3_xboole_0(X2,X3))=k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_47, negated_conjecture, (k2_zfmisc_1(esk6_0,esk7_0)=k2_zfmisc_1(esk8_0,esk9_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_48, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_49, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),X3)=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_51, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_52, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_53, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_45, c_0_28])).
cnf(c_0_54, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk8_0,esk9_0),k2_zfmisc_1(esk6_0,X1))=k2_zfmisc_1(esk6_0,k3_xboole_0(X1,esk7_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_28])).
fof(c_0_55, plain, ![X50, X51]:(~r1_tarski(X50,X51)|k3_xboole_0(X50,X51)=X50), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_56, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|r2_hidden(esk4_2(k4_xboole_0(X1,X2),X3),X1)), inference(spm,[status(thm)],[c_0_48, c_0_29])).
cnf(c_0_57, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk8_0,esk9_0),k2_zfmisc_1(X1,esk7_0))=k2_zfmisc_1(k3_xboole_0(esk6_0,X1),esk7_0)), inference(spm,[status(thm)],[c_0_49, c_0_47])).
cnf(c_0_58, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(k4_xboole_0(X3,X1),X4))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52])).
cnf(c_0_59, plain, (r1_tarski(k2_zfmisc_1(X1,k3_xboole_0(X2,X3)),k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_53, c_0_46])).
cnf(c_0_60, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk8_0,esk6_0),esk9_0)=k2_zfmisc_1(esk6_0,k3_xboole_0(esk9_0,esk7_0))), inference(spm,[status(thm)],[c_0_54, c_0_49])).
fof(c_0_61, plain, ![X52]:k3_xboole_0(X52,X52)=X52, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_62, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_63, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_36, c_0_56])).
cnf(c_0_64, plain, (X1=k1_xboole_0|X2=k1_xboole_0|k2_zfmisc_1(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_65, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk6_0,k4_xboole_0(X1,esk8_0)),esk7_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_66, negated_conjecture, (esk7_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_67, plain, ![X53, X54]:(((r1_tarski(X53,X54)|X53!=X54)&(r1_tarski(X54,X53)|X53!=X54))&(~r1_tarski(X53,X54)|~r1_tarski(X54,X53)|X53=X54)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_68, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk6_0,k3_xboole_0(X1,esk7_0)),k2_zfmisc_1(esk6_0,k3_xboole_0(esk9_0,esk7_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_50]), c_0_54])).
cnf(c_0_69, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_61])).
fof(c_0_70, plain, ![X55, X56]:((k4_xboole_0(X55,X56)!=k1_xboole_0|r1_tarski(X55,X56))&(~r1_tarski(X55,X56)|k4_xboole_0(X55,X56)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_71, plain, (k3_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_28])).
cnf(c_0_72, negated_conjecture, (k3_xboole_0(esk6_0,k4_xboole_0(X1,esk8_0))=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_66])).
fof(c_0_73, plain, ![X57, X58, X60, X61, X62]:((r1_xboole_0(X57,X58)|r2_hidden(esk10_2(X57,X58),k3_xboole_0(X57,X58)))&(~r2_hidden(X62,k3_xboole_0(X60,X61))|~r1_xboole_0(X60,X61))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
cnf(c_0_74, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_75, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk8_0,esk9_0),k2_zfmisc_1(esk6_0,k3_xboole_0(esk9_0,esk7_0)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_47])).
cnf(c_0_76, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk6_0,k3_xboole_0(X1,esk7_0)),k2_zfmisc_1(esk8_0,esk9_0))), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_77, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_78, negated_conjecture, (k4_xboole_0(esk6_0,esk8_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_79, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_73])).
fof(c_0_80, plain, ![X63]:r1_xboole_0(X63,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_81, negated_conjecture, (k2_zfmisc_1(esk6_0,k3_xboole_0(esk9_0,esk7_0))=k2_zfmisc_1(esk8_0,esk9_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76])])).
cnf(c_0_82, negated_conjecture, (r1_tarski(esk6_0,esk8_0)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_83, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk10_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_84, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_85, plain, (~r1_xboole_0(X1,X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_79, c_0_69])).
cnf(c_0_86, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_87, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk8_0,esk6_0),esk9_0)=k2_zfmisc_1(esk8_0,esk9_0)), inference(rw,[status(thm)],[c_0_60, c_0_81])).
cnf(c_0_88, negated_conjecture, (k3_xboole_0(esk8_0,esk6_0)=esk6_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_82]), c_0_28])).
cnf(c_0_89, plain, (~r1_xboole_0(X1,X2)|~r2_hidden(X3,k3_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_79, c_0_28])).
cnf(c_0_90, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk10_2(X1,X2),k3_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_83, c_0_28])).
cnf(c_0_91, plain, (r1_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))|r2_hidden(esk10_2(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)),k2_zfmisc_1(X1,k3_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_83, c_0_46])).
cnf(c_0_92, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_84])).
cnf(c_0_93, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_94, negated_conjecture, (k2_zfmisc_1(esk6_0,esk9_0)=k2_zfmisc_1(esk8_0,esk9_0)), inference(rw,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_95, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_89, c_0_22])).
cnf(c_0_96, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_79, c_0_90])).
cnf(c_0_97, plain, (r1_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,k4_xboole_0(X3,X2)))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_51]), c_0_92]), c_0_93])).
cnf(c_0_98, negated_conjecture, (k2_zfmisc_1(esk6_0,k3_xboole_0(X1,esk7_0))=k2_zfmisc_1(esk6_0,k3_xboole_0(X1,esk9_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_94]), c_0_28]), c_0_54])).
cnf(c_0_99, negated_conjecture, (esk6_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_100, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(esk6_0,k3_xboole_0(X1,esk7_0)))|~r1_xboole_0(k2_zfmisc_1(esk6_0,X1),k2_zfmisc_1(esk8_0,esk9_0))), inference(spm,[status(thm)],[c_0_95, c_0_54])).
cnf(c_0_101, plain, (r1_xboole_0(k2_zfmisc_1(X1,k4_xboole_0(X2,X3)),k2_zfmisc_1(X1,X3))), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_102, negated_conjecture, (k3_xboole_0(X1,esk7_0)=k1_xboole_0|k2_zfmisc_1(esk6_0,k3_xboole_0(X1,esk9_0))!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_98]), c_0_99])).
cnf(c_0_103, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(esk6_0,k3_xboole_0(X1,esk9_0)))|~r1_xboole_0(k2_zfmisc_1(esk6_0,X1),k2_zfmisc_1(esk8_0,esk9_0))), inference(rw,[status(thm)],[c_0_100, c_0_98])).
cnf(c_0_104, negated_conjecture, (r1_xboole_0(k2_zfmisc_1(esk6_0,k4_xboole_0(X1,esk7_0)),k2_zfmisc_1(esk8_0,esk9_0))), inference(spm,[status(thm)],[c_0_101, c_0_47])).
cnf(c_0_105, negated_conjecture, (k3_xboole_0(X1,esk7_0)=k1_xboole_0|k2_zfmisc_1(esk6_0,k3_xboole_0(esk9_0,X1))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_102, c_0_28])).
cnf(c_0_106, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(esk6_0,k3_xboole_0(esk9_0,k4_xboole_0(X1,esk7_0))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_28])).
cnf(c_0_107, negated_conjecture, (k3_xboole_0(esk7_0,k4_xboole_0(X1,esk9_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_51]), c_0_92])]), c_0_28])).
cnf(c_0_108, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(esk6_0,k4_xboole_0(esk9_0,esk7_0)))), inference(spm,[status(thm)],[c_0_106, c_0_71])).
cnf(c_0_109, negated_conjecture, (k4_xboole_0(esk7_0,esk9_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_71, c_0_107])).
cnf(c_0_110, negated_conjecture, (k2_zfmisc_1(esk6_0,k4_xboole_0(esk9_0,esk7_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_42, c_0_108])).
cnf(c_0_111, negated_conjecture, (r1_tarski(esk7_0,esk9_0)), inference(spm,[status(thm)],[c_0_77, c_0_109])).
cnf(c_0_112, negated_conjecture, (k4_xboole_0(esk9_0,esk7_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_110]), c_0_99])).
cnf(c_0_113, negated_conjecture, (esk7_0=esk9_0|~r1_tarski(esk9_0,esk7_0)), inference(spm,[status(thm)],[c_0_74, c_0_111])).
cnf(c_0_114, negated_conjecture, (r1_tarski(esk9_0,esk7_0)), inference(spm,[status(thm)],[c_0_77, c_0_112])).
cnf(c_0_115, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk6_0,X1),esk9_0)=k2_zfmisc_1(k3_xboole_0(esk8_0,X1),esk9_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_94]), c_0_49])).
cnf(c_0_116, negated_conjecture, (esk7_0=esk9_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_113, c_0_114])])).
cnf(c_0_117, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk8_0,k4_xboole_0(X1,esk6_0)),esk9_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_51]), c_0_52])).
cnf(c_0_118, negated_conjecture, (esk9_0!=k1_xboole_0), inference(rw,[status(thm)],[c_0_66, c_0_116])).
cnf(c_0_119, negated_conjecture, (k3_xboole_0(esk8_0,k4_xboole_0(X1,esk6_0))=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_117]), c_0_118])).
cnf(c_0_120, negated_conjecture, (k4_xboole_0(esk8_0,esk6_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_71, c_0_119])).
cnf(c_0_121, negated_conjecture, (esk6_0!=esk8_0|esk7_0!=esk9_0), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_122, negated_conjecture, (esk6_0=esk8_0|~r1_tarski(esk8_0,esk6_0)), inference(spm,[status(thm)],[c_0_74, c_0_82])).
cnf(c_0_123, negated_conjecture, (r1_tarski(esk8_0,esk6_0)), inference(spm,[status(thm)],[c_0_77, c_0_120])).
cnf(c_0_124, negated_conjecture, (esk6_0!=esk8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_121, c_0_116])])).
cnf(c_0_125, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_122, c_0_123])]), c_0_124]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 126
# Proof object clause steps            : 93
# Proof object formula steps           : 33
# Proof object conjectures             : 46
# Proof object clause conjectures      : 43
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 55
# Proof object simplifying inferences  : 39
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 39
# Processed clauses                    : 2478
# ...of these trivial                  : 76
# ...subsumed                          : 1869
# ...remaining for further processing  : 533
# Other redundant clauses eliminated   : 11
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 21
# Backward-rewritten                   : 161
# Generated clauses                    : 15605
# ...of the previous two non-trivial   : 9774
# Contextual simplify-reflections      : 2
# Paramodulations                      : 15588
# Factorizations                       : 6
# Equation resolutions                 : 11
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 341
#    Positive orientable unit clauses  : 100
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 13
#    Non-unit-clauses                  : 227
# Current number of unprocessed clauses: 7205
# ...number of literals in the above   : 15695
# Current number of archived formulas  : 0
# Current number of archived clauses   : 182
# Clause-clause subsumption calls (NU) : 28069
# Rec. Clause-clause subsumption calls : 22710
# Non-unit clause-clause subsumptions  : 892
# Unit Clause-clause subsumption calls : 4125
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 99
# BW rewrite match successes           : 37
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 151267

# -------------------------------------------------
# User time                : 0.083 s
# System time              : 0.010 s
# Total time               : 0.092 s
# Maximum resident set size: 28648 pages
