# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C41_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.007 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t19_setfam_1, conjecture, ![X1, X2]:(r2_setfam_1(X2,X1)=>(X1=k1_xboole_0|r1_tarski(k1_setfam_1(X2),k1_setfam_1(X1)))), file('data/MPTP/proofs/78690cd3a4b14a26bcaf1eda71b5deaf.E_input', t19_setfam_1)).
fof(d3_setfam_1, axiom, ![X1, X2]:(r2_setfam_1(X1,X2)<=>![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X1)&r1_tarski(X4,X3)))))), file('data/MPTP/proofs/78690cd3a4b14a26bcaf1eda71b5deaf.E_input', d3_setfam_1)).
fof(t6_setfam_1, axiom, ![X1, X2]:(![X3]:(r2_hidden(X3,X1)=>r1_tarski(X2,X3))=>(X1=k1_xboole_0|r1_tarski(X2,k1_setfam_1(X1)))), file('data/MPTP/proofs/78690cd3a4b14a26bcaf1eda71b5deaf.E_input', t6_setfam_1)).
fof(t8_setfam_1, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&r1_tarski(X1,X3))=>r1_tarski(k1_setfam_1(X2),X3)), file('data/MPTP/proofs/78690cd3a4b14a26bcaf1eda71b5deaf.E_input', t8_setfam_1)).
fof(c_0_4, negated_conjecture, ~(![X1, X2]:(r2_setfam_1(X2,X1)=>(X1=k1_xboole_0|r1_tarski(k1_setfam_1(X2),k1_setfam_1(X1))))), inference(assume_negation,[status(cth)],[t19_setfam_1])).
fof(c_0_5, plain, ![X7, X8, X9, X11, X12, X14]:(((r2_hidden(esk3_3(X7,X8,X9),X7)|~r2_hidden(X9,X8)|~r2_setfam_1(X7,X8))&(r1_tarski(esk3_3(X7,X8,X9),X9)|~r2_hidden(X9,X8)|~r2_setfam_1(X7,X8)))&((r2_hidden(esk4_2(X11,X12),X12)|r2_setfam_1(X11,X12))&(~r2_hidden(X14,X11)|~r1_tarski(X14,esk4_2(X11,X12))|r2_setfam_1(X11,X12)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_setfam_1])])])])])])).
fof(c_0_6, negated_conjecture, (r2_setfam_1(esk2_0,esk1_0)&(esk1_0!=k1_xboole_0&~r1_tarski(k1_setfam_1(esk2_0),k1_setfam_1(esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
cnf(c_0_7, plain, (r2_hidden(esk3_3(X1,X2,X3),X1)|~r2_hidden(X3,X2)|~r2_setfam_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8, negated_conjecture, (r2_setfam_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_9, plain, ![X15, X16]:((r2_hidden(esk5_2(X15,X16),X15)|(X15=k1_xboole_0|r1_tarski(X16,k1_setfam_1(X15))))&(~r1_tarski(X16,esk5_2(X15,X16))|(X15=k1_xboole_0|r1_tarski(X16,k1_setfam_1(X15))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_setfam_1])])])])).
fof(c_0_10, plain, ![X18, X19, X20]:(~r2_hidden(X18,X19)|~r1_tarski(X18,X20)|r1_tarski(k1_setfam_1(X19),X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_setfam_1])])).
cnf(c_0_11, negated_conjecture, (r2_hidden(esk3_3(esk2_0,esk1_0,X1),esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_12, plain, (r2_hidden(esk5_2(X1,X2),X1)|X1=k1_xboole_0|r1_tarski(X2,k1_setfam_1(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14, plain, (r1_tarski(esk3_3(X1,X2,X3),X3)|~r2_hidden(X3,X2)|~r2_setfam_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_15, plain, (r1_tarski(k1_setfam_1(X2),X3)|~r2_hidden(X1,X2)|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, negated_conjecture, (r2_hidden(esk3_3(esk2_0,esk1_0,esk5_2(esk1_0,X1)),esk2_0)|r1_tarski(X1,k1_setfam_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])).
cnf(c_0_17, plain, (X1=k1_xboole_0|r1_tarski(esk3_3(X2,X1,esk5_2(X1,X3)),esk5_2(X1,X3))|r1_tarski(X3,k1_setfam_1(X1))|~r2_setfam_1(X2,X1)), inference(spm,[status(thm)],[c_0_14, c_0_12])).
cnf(c_0_18, negated_conjecture, (r1_tarski(X1,k1_setfam_1(esk1_0))|r1_tarski(k1_setfam_1(esk2_0),X2)|~r1_tarski(esk3_3(esk2_0,esk1_0,esk5_2(esk1_0,X1)),X2)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_19, negated_conjecture, (r1_tarski(esk3_3(esk2_0,esk1_0,esk5_2(esk1_0,X1)),esk5_2(esk1_0,X1))|r1_tarski(X1,k1_setfam_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_8]), c_0_13])).
cnf(c_0_20, plain, (X2=k1_xboole_0|r1_tarski(X1,k1_setfam_1(X2))|~r1_tarski(X1,esk5_2(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_21, negated_conjecture, (r1_tarski(k1_setfam_1(esk2_0),esk5_2(esk1_0,X1))|r1_tarski(X1,k1_setfam_1(esk1_0))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_22, negated_conjecture, (~r1_tarski(k1_setfam_1(esk2_0),k1_setfam_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_23, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_13]), c_0_22]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 24
# Proof object clause steps            : 15
# Proof object formula steps           : 9
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 4
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 10
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 10
# Processed clauses                    : 27
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 27
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 19
# ...of the previous two non-trivial   : 17
# Contextual simplify-reflections      : 0
# Paramodulations                      : 19
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 27
#    Positive orientable unit clauses  : 1
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 24
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 0
# Clause-clause subsumption calls (NU) : 16
# Rec. Clause-clause subsumption calls : 14
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1109

# -------------------------------------------------
# User time                : 0.005 s
# System time              : 0.003 s
# Total time               : 0.008 s
# Maximum resident set size: 28648 pages
