# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t31_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>r1_tarski(k3_relat_1(X1),k3_relat_1(X2))))), file('data/MPTP/proofs/7898428f304b40aab2be1f463390a3b9.E_input', t31_relat_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/7898428f304b40aab2be1f463390a3b9.E_input', d3_tarski)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/7898428f304b40aab2be1f463390a3b9.E_input', d5_relat_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/7898428f304b40aab2be1f463390a3b9.E_input', t7_xboole_1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/7898428f304b40aab2be1f463390a3b9.E_input', l51_zfmisc_1)).
fof(d6_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k3_relat_1(X1)=k2_xboole_0(k1_relat_1(X1),k2_relat_1(X1))), file('data/MPTP/proofs/7898428f304b40aab2be1f463390a3b9.E_input', d6_relat_1)).
fof(t30_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)=>(r2_hidden(X1,k3_relat_1(X3))&r2_hidden(X2,k3_relat_1(X3))))), file('data/MPTP/proofs/7898428f304b40aab2be1f463390a3b9.E_input', t30_relat_1)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/7898428f304b40aab2be1f463390a3b9.E_input', cc2_relat_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/7898428f304b40aab2be1f463390a3b9.E_input', t3_subset)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/7898428f304b40aab2be1f463390a3b9.E_input', t1_xboole_1)).
fof(t25_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>(r1_tarski(k1_relat_1(X1),k1_relat_1(X2))&r1_tarski(k2_relat_1(X1),k2_relat_1(X2)))))), file('data/MPTP/proofs/7898428f304b40aab2be1f463390a3b9.E_input', t25_relat_1)).
fof(t8_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X3,X2))=>r1_tarski(k2_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/7898428f304b40aab2be1f463390a3b9.E_input', t8_xboole_1)).
fof(c_0_12, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>r1_tarski(k3_relat_1(X1),k3_relat_1(X2)))))), inference(assume_negation,[status(cth)],[t31_relat_1])).
fof(c_0_13, plain, ![X18, X19, X20, X21, X22]:((~r1_tarski(X18,X19)|(~r2_hidden(X20,X18)|r2_hidden(X20,X19)))&((r2_hidden(esk3_2(X21,X22),X21)|r1_tarski(X21,X22))&(~r2_hidden(esk3_2(X21,X22),X22)|r1_tarski(X21,X22)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_14, negated_conjecture, (v1_relat_1(esk1_0)&(v1_relat_1(esk2_0)&(r1_tarski(esk1_0,esk2_0)&~r1_tarski(k3_relat_1(esk1_0),k3_relat_1(esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
fof(c_0_15, plain, ![X26, X27, X28, X30, X31, X32, X33, X35]:(((~r2_hidden(X28,X27)|r2_hidden(k4_tarski(esk4_3(X26,X27,X28),X28),X26)|X27!=k2_relat_1(X26))&(~r2_hidden(k4_tarski(X31,X30),X26)|r2_hidden(X30,X27)|X27!=k2_relat_1(X26)))&((~r2_hidden(esk5_2(X32,X33),X33)|~r2_hidden(k4_tarski(X35,esk5_2(X32,X33)),X32)|X33=k2_relat_1(X32))&(r2_hidden(esk5_2(X32,X33),X33)|r2_hidden(k4_tarski(esk6_2(X32,X33),esk5_2(X32,X33)),X32)|X33=k2_relat_1(X32)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_16, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (r2_hidden(k4_tarski(esk4_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_19, plain, ![X24, X25]:r1_tarski(X24,k2_xboole_0(X24,X25)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_20, plain, ![X7, X8]:k3_tarski(k2_tarski(X7,X8))=k2_xboole_0(X7,X8), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_21, plain, ![X12]:(~v1_relat_1(X12)|k3_relat_1(X12)=k2_xboole_0(k1_relat_1(X12),k2_relat_1(X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_relat_1])])).
fof(c_0_22, plain, ![X15, X16, X17]:((r2_hidden(X15,k3_relat_1(X17))|~r2_hidden(k4_tarski(X15,X16),X17)|~v1_relat_1(X17))&(r2_hidden(X16,k3_relat_1(X17))|~r2_hidden(k4_tarski(X15,X16),X17)|~v1_relat_1(X17))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_relat_1])])])).
cnf(c_0_23, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_24, plain, (r2_hidden(k4_tarski(esk4_3(X1,k2_relat_1(X1),X2),X2),X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (k3_relat_1(X1)=k2_xboole_0(k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_28, plain, ![X13, X14]:(~v1_relat_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(X13))|v1_relat_1(X14))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_29, plain, ![X37, X38]:((~m1_subset_1(X37,k1_zfmisc_1(X38))|r1_tarski(X37,X38))&(~r1_tarski(X37,X38)|m1_subset_1(X37,k1_zfmisc_1(X38)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_30, plain, (r2_hidden(X1,k3_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, negated_conjecture, (r2_hidden(k4_tarski(esk4_3(esk1_0,k2_relat_1(esk1_0),X1),X1),esk2_0)|~r2_hidden(X1,k2_relat_1(esk1_0))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_32, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_33, plain, ![X41, X42, X43]:(~r1_tarski(X41,X42)|~r1_tarski(X42,X43)|r1_tarski(X41,X43)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
cnf(c_0_34, plain, (r1_tarski(X1,k3_tarski(k2_tarski(X1,X2)))), inference(rw,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_35, plain, (k3_relat_1(X1)=k3_tarski(k2_tarski(k1_relat_1(X1),k2_relat_1(X1)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_27, c_0_26])).
fof(c_0_36, plain, ![X39, X40]:((r1_tarski(k1_relat_1(X39),k1_relat_1(X40))|~r1_tarski(X39,X40)|~v1_relat_1(X40)|~v1_relat_1(X39))&(r1_tarski(k2_relat_1(X39),k2_relat_1(X40))|~r1_tarski(X39,X40)|~v1_relat_1(X40)|~v1_relat_1(X39))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t25_relat_1])])])])).
cnf(c_0_37, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_39, plain, ![X9, X10, X11]:(~r1_tarski(X9,X10)|~r1_tarski(X11,X10)|r1_tarski(k2_xboole_0(X9,X11),X10)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_xboole_1])])).
cnf(c_0_40, plain, (r1_tarski(X1,X2)|~r2_hidden(esk3_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_41, negated_conjecture, (r2_hidden(X1,k3_relat_1(esk2_0))|~r2_hidden(X1,k2_relat_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])])).
cnf(c_0_42, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_43, plain, (r1_tarski(k1_relat_1(X1),k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_44, plain, (r1_tarski(k1_relat_1(X1),k1_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_45, plain, (v1_relat_1(X1)|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_46, plain, (r1_tarski(k2_xboole_0(X1,X3),X2)|~r1_tarski(X1,X2)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_47, negated_conjecture, (r1_tarski(X1,k3_relat_1(esk2_0))|~r2_hidden(esk3_2(X1,k3_relat_1(esk2_0)),k2_relat_1(esk1_0))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_48, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_49, plain, (r1_tarski(X1,k3_relat_1(X2))|~r1_tarski(X1,k1_relat_1(X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_50, plain, (r1_tarski(k1_relat_1(X1),k1_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(csr,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_51, negated_conjecture, (r1_tarski(X1,esk2_0)|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_42, c_0_17])).
cnf(c_0_52, plain, (r1_tarski(k3_tarski(k2_tarski(X1,X3)),X2)|~r1_tarski(X3,X2)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_46, c_0_26])).
cnf(c_0_53, negated_conjecture, (r1_tarski(k2_relat_1(esk1_0),k3_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_54, plain, (r1_tarski(k1_relat_1(X1),k3_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_55, negated_conjecture, (r1_tarski(k3_tarski(k2_tarski(X1,X2)),esk2_0)|~r1_tarski(X2,esk1_0)|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_56, negated_conjecture, (r1_tarski(X1,k3_relat_1(esk2_0))|~r1_tarski(X1,k2_relat_1(esk1_0))), inference(spm,[status(thm)],[c_0_42, c_0_53])).
cnf(c_0_57, plain, (r1_tarski(k2_relat_1(X1),k2_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_58, plain, (r1_tarski(X1,k3_relat_1(X2))|~r1_tarski(X1,k1_relat_1(X3))|~r1_tarski(X3,X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_42, c_0_54])).
cnf(c_0_59, plain, (v1_relat_1(X1)|~v1_relat_1(k3_tarski(k2_tarski(X1,X2)))), inference(spm,[status(thm)],[c_0_45, c_0_34])).
cnf(c_0_60, negated_conjecture, (v1_relat_1(k3_tarski(k2_tarski(X1,X2)))|~r1_tarski(X2,esk1_0)|~r1_tarski(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_55]), c_0_32])])).
cnf(c_0_61, negated_conjecture, (r1_tarski(X1,k3_relat_1(esk2_0))|~r1_tarski(X2,k2_relat_1(esk1_0))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_42, c_0_56])).
cnf(c_0_62, plain, (r1_tarski(k2_relat_1(X1),k2_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)), inference(csr,[status(thm)],[c_0_57, c_0_45])).
cnf(c_0_63, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_64, plain, (r1_tarski(k1_relat_1(X1),k3_relat_1(X2))|~r1_tarski(X3,X2)|~r1_tarski(X1,X3)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_50]), c_0_45])).
cnf(c_0_65, negated_conjecture, (v1_relat_1(X1)|~r1_tarski(X2,esk1_0)|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_66, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_40, c_0_48])).
cnf(c_0_67, negated_conjecture, (r1_tarski(X1,k3_relat_1(esk2_0))|~r1_tarski(X1,k2_relat_1(X2))|~r1_tarski(X2,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63])])).
cnf(c_0_68, plain, (r1_tarski(k3_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X2)|~r1_tarski(k1_relat_1(X1),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_52, c_0_35])).
cnf(c_0_69, negated_conjecture, (r1_tarski(k1_relat_1(X1),k3_relat_1(esk2_0))|~r1_tarski(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_17]), c_0_32])])).
cnf(c_0_70, negated_conjecture, (v1_relat_1(X1)|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_71, negated_conjecture, (r1_tarski(k2_relat_1(X1),k3_relat_1(esk2_0))|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_67, c_0_66])).
cnf(c_0_72, negated_conjecture, (~r1_tarski(k3_relat_1(esk1_0),k3_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_73, negated_conjecture, (r1_tarski(k3_relat_1(X1),k3_relat_1(esk2_0))|~r1_tarski(X1,esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70]), c_0_71])).
cnf(c_0_74, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_66])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 75
# Proof object clause steps            : 50
# Proof object formula steps           : 25
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 26
# Proof object simplifying inferences  : 19
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 22
# Processed clauses                    : 289
# ...of these trivial                  : 0
# ...subsumed                          : 103
# ...remaining for further processing  : 186
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 0
# Generated clauses                    : 571
# ...of the previous two non-trivial   : 538
# Contextual simplify-reflections      : 5
# Paramodulations                      : 567
# Factorizations                       : 2
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 158
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 150
# Current number of unprocessed clauses: 263
# ...number of literals in the above   : 1028
# Current number of archived formulas  : 0
# Current number of archived clauses   : 27
# Clause-clause subsumption calls (NU) : 5366
# Rec. Clause-clause subsumption calls : 3521
# Non-unit clause-clause subsumptions  : 112
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 10377

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.003 s
# Total time               : 0.018 s
# Maximum resident set size: 28788 pages
