# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_filter_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_partfun1(X2,X1)&(v3_relat_2(X2)&(v8_relat_2(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>![X3]:(m1_subset_1(X3,X1)=>![X4]:(m2_filter_1(X4,X1,X2)=>(r3_binop_1(X1,X3,X4)=>r3_binop_1(k8_eqrel_1(X1,X2),k9_eqrel_1(X1,X2,X3),k3_filter_1(X1,X2,X4))))))), file('data/MPTP/proofs/78e76444200e43708bcdc97cde727105.E_input', t8_filter_1)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/78e76444200e43708bcdc97cde727105.E_input', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/78e76444200e43708bcdc97cde727105.E_input', fc6_relat_1)).
fof(dt_m2_filter_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&v1_relat_1(X2))=>![X3]:(m2_filter_1(X3,X1,X2)=>(v1_funct_1(X3)&(v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))))), file('data/MPTP/proofs/78e76444200e43708bcdc97cde727105.E_input', dt_m2_filter_1)).
fof(d7_binop_1, axiom, ![X1, X2]:(m1_subset_1(X2,X1)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))=>(r3_binop_1(X1,X2,X3)<=>(r1_binop_1(X1,X2,X3)&r2_binop_1(X1,X2,X3))))), file('data/MPTP/proofs/78e76444200e43708bcdc97cde727105.E_input', d7_binop_1)).
fof(dt_k9_eqrel_1, axiom, ![X1, X2, X3]:((~(v1_xboole_0(X1))&((v3_relat_2(X2)&(v8_relat_2(X2)&(v1_partfun1(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))&m1_subset_1(X3,X1)))=>m2_subset_1(k9_eqrel_1(X1,X2,X3),k1_zfmisc_1(X1),k8_eqrel_1(X1,X2))), file('data/MPTP/proofs/78e76444200e43708bcdc97cde727105.E_input', dt_k9_eqrel_1)).
fof(dt_k8_eqrel_1, axiom, ![X1, X2]:((v3_relat_2(X2)&(v8_relat_2(X2)&(v1_partfun1(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>m1_eqrel_1(k8_eqrel_1(X1,X2),X1)), file('data/MPTP/proofs/78e76444200e43708bcdc97cde727105.E_input', dt_k8_eqrel_1)).
fof(fc3_eqrel_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&(v3_relat_2(X2)&(v8_relat_2(X2)&(v1_partfun1(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))))=>~(v1_xboole_0(k7_eqrel_1(X1,X2)))), file('data/MPTP/proofs/78e76444200e43708bcdc97cde727105.E_input', fc3_eqrel_1)).
fof(redefinition_k8_eqrel_1, axiom, ![X1, X2]:((v3_relat_2(X2)&(v8_relat_2(X2)&(v1_partfun1(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>k8_eqrel_1(X1,X2)=k7_eqrel_1(X1,X2)), file('data/MPTP/proofs/78e76444200e43708bcdc97cde727105.E_input', redefinition_k8_eqrel_1)).
fof(t7_filter_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_partfun1(X2,X1)&(v3_relat_2(X2)&(v8_relat_2(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>![X3]:(m1_subset_1(X3,X1)=>![X4]:(m2_filter_1(X4,X1,X2)=>(r2_binop_1(X1,X3,X4)=>r2_binop_1(k8_eqrel_1(X1,X2),k9_eqrel_1(X1,X2,X3),k3_filter_1(X1,X2,X4))))))), file('data/MPTP/proofs/78e76444200e43708bcdc97cde727105.E_input', t7_filter_1)).
fof(t6_filter_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_partfun1(X2,X1)&(v3_relat_2(X2)&(v8_relat_2(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>![X3]:(m1_subset_1(X3,X1)=>![X4]:(m2_filter_1(X4,X1,X2)=>(r1_binop_1(X1,X3,X4)=>r1_binop_1(k8_eqrel_1(X1,X2),k9_eqrel_1(X1,X2,X3),k3_filter_1(X1,X2,X4))))))), file('data/MPTP/proofs/78e76444200e43708bcdc97cde727105.E_input', t6_filter_1)).
fof(dt_k3_filter_1, axiom, ![X1, X2, X3]:((~(v1_xboole_0(X1))&((v1_partfun1(X2,X1)&(v3_relat_2(X2)&(v8_relat_2(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))&(v1_funct_1(X3)&(v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))))=>(v1_funct_1(k3_filter_1(X1,X2,X3))&(v1_funct_2(k3_filter_1(X1,X2,X3),k2_zfmisc_1(k8_eqrel_1(X1,X2),k8_eqrel_1(X1,X2)),k8_eqrel_1(X1,X2))&m1_subset_1(k3_filter_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k8_eqrel_1(X1,X2),k8_eqrel_1(X1,X2)),k8_eqrel_1(X1,X2))))))), file('data/MPTP/proofs/78e76444200e43708bcdc97cde727105.E_input', dt_k3_filter_1)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&(~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(X1))))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('data/MPTP/proofs/78e76444200e43708bcdc97cde727105.E_input', redefinition_m2_subset_1)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('data/MPTP/proofs/78e76444200e43708bcdc97cde727105.E_input', cc1_subset_1)).
fof(dt_m1_eqrel_1, axiom, ![X1, X2]:(m1_eqrel_1(X2,X1)=>m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), file('data/MPTP/proofs/78e76444200e43708bcdc97cde727105.E_input', dt_m1_eqrel_1)).
fof(c_0_15, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_partfun1(X2,X1)&(v3_relat_2(X2)&(v8_relat_2(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>![X3]:(m1_subset_1(X3,X1)=>![X4]:(m2_filter_1(X4,X1,X2)=>(r3_binop_1(X1,X3,X4)=>r3_binop_1(k8_eqrel_1(X1,X2),k9_eqrel_1(X1,X2,X3),k3_filter_1(X1,X2,X4)))))))), inference(assume_negation,[status(cth)],[t8_filter_1])).
fof(c_0_16, plain, ![X9, X10]:(~v1_relat_1(X9)|(~m1_subset_1(X10,k1_zfmisc_1(X9))|v1_relat_1(X10))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_17, negated_conjecture, (~v1_xboole_0(esk1_0)&((v1_partfun1(esk2_0,esk1_0)&(v3_relat_2(esk2_0)&(v8_relat_2(esk2_0)&m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0))))))&(m1_subset_1(esk3_0,esk1_0)&(m2_filter_1(esk4_0,esk1_0,esk2_0)&(r3_binop_1(esk1_0,esk3_0,esk4_0)&~r3_binop_1(k8_eqrel_1(esk1_0,esk2_0),k9_eqrel_1(esk1_0,esk2_0,esk3_0),k3_filter_1(esk1_0,esk2_0,esk4_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_15])])])])).
fof(c_0_18, plain, ![X11, X12]:v1_relat_1(k2_zfmisc_1(X11,X12)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_19, plain, ![X13, X14, X15]:((v1_funct_1(X15)|~m2_filter_1(X15,X13,X14)|(v1_xboole_0(X13)|~v1_relat_1(X14)))&((v1_funct_2(X15,k2_zfmisc_1(X13,X13),X13)|~m2_filter_1(X15,X13,X14)|(v1_xboole_0(X13)|~v1_relat_1(X14)))&(m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X13,X13),X13)))|~m2_filter_1(X15,X13,X14)|(v1_xboole_0(X13)|~v1_relat_1(X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_filter_1])])])])])).
cnf(c_0_20, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, plain, (v1_funct_1(X1)|v1_xboole_0(X2)|~m2_filter_1(X1,X2,X3)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, negated_conjecture, (m2_filter_1(esk4_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_26, plain, ![X16, X17, X18]:(((r1_binop_1(X16,X17,X18)|~r3_binop_1(X16,X17,X18)|(~v1_funct_1(X18)|(~v1_funct_2(X18,k2_zfmisc_1(X16,X16),X16)|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X16,X16),X16)))))|~m1_subset_1(X17,X16))&(r2_binop_1(X16,X17,X18)|~r3_binop_1(X16,X17,X18)|(~v1_funct_1(X18)|(~v1_funct_2(X18,k2_zfmisc_1(X16,X16),X16)|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X16,X16),X16)))))|~m1_subset_1(X17,X16)))&(~r1_binop_1(X16,X17,X18)|~r2_binop_1(X16,X17,X18)|r3_binop_1(X16,X17,X18)|(~v1_funct_1(X18)|(~v1_funct_2(X18,k2_zfmisc_1(X16,X16),X16)|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X16,X16),X16)))))|~m1_subset_1(X17,X16))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_binop_1])])])])).
cnf(c_0_27, plain, (v1_funct_2(X1,k2_zfmisc_1(X2,X2),X2)|v1_xboole_0(X2)|~m2_filter_1(X1,X2,X3)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, negated_conjecture, (v1_relat_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])])).
cnf(c_0_29, negated_conjecture, (v1_funct_1(esk4_0)|~v1_relat_1(esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
cnf(c_0_30, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|v1_xboole_0(X2)|~m2_filter_1(X1,X2,X3)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_31, plain, ![X19, X20, X21]:(v1_xboole_0(X19)|(~v3_relat_2(X20)|(~v8_relat_2(X20)|(~v1_partfun1(X20,X19)|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(X19,X19)))))|~m1_subset_1(X21,X19))|m2_subset_1(k9_eqrel_1(X19,X20,X21),k1_zfmisc_1(X19),k8_eqrel_1(X19,X20))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k9_eqrel_1])])])).
fof(c_0_32, plain, ![X22, X23]:(~v3_relat_2(X23)|(~v8_relat_2(X23)|(~v1_partfun1(X23,X22)|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(X22,X22)))))|m1_eqrel_1(k8_eqrel_1(X22,X23),X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_eqrel_1])])).
fof(c_0_33, plain, ![X24, X25]:(v1_xboole_0(X24)|(~v3_relat_2(X25)|(~v8_relat_2(X25)|(~v1_partfun1(X25,X24)|~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(X24,X24))))))|~v1_xboole_0(k7_eqrel_1(X24,X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc3_eqrel_1])])])).
fof(c_0_34, plain, ![X26, X27]:(~v3_relat_2(X27)|(~v8_relat_2(X27)|(~v1_partfun1(X27,X26)|~m1_subset_1(X27,k1_zfmisc_1(k2_zfmisc_1(X26,X26)))))|k8_eqrel_1(X26,X27)=k7_eqrel_1(X26,X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_eqrel_1])])).
fof(c_0_35, plain, ![X32, X33, X34, X35]:(v1_xboole_0(X32)|(~v1_partfun1(X33,X32)|(~v3_relat_2(X33)|(~v8_relat_2(X33)|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X32,X32)))))|(~m1_subset_1(X34,X32)|(~m2_filter_1(X35,X32,X33)|(~r2_binop_1(X32,X34,X35)|r2_binop_1(k8_eqrel_1(X32,X33),k9_eqrel_1(X32,X33,X34),k3_filter_1(X32,X33,X35))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_filter_1])])])])).
cnf(c_0_36, plain, (r2_binop_1(X1,X2,X3)|~r3_binop_1(X1,X2,X3)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_37, negated_conjecture, (r3_binop_1(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_38, negated_conjecture, (v1_funct_2(esk4_0,k2_zfmisc_1(esk1_0,esk1_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_24]), c_0_28])]), c_0_25])).
cnf(c_0_39, negated_conjecture, (v1_funct_1(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_28])])).
cnf(c_0_40, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0),esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_24]), c_0_28])]), c_0_25])).
cnf(c_0_41, negated_conjecture, (m1_subset_1(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_42, plain, ![X28, X29, X30, X31]:(v1_xboole_0(X28)|(~v1_partfun1(X29,X28)|(~v3_relat_2(X29)|(~v8_relat_2(X29)|~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(X28,X28)))))|(~m1_subset_1(X30,X28)|(~m2_filter_1(X31,X28,X29)|(~r1_binop_1(X28,X30,X31)|r1_binop_1(k8_eqrel_1(X28,X29),k9_eqrel_1(X28,X29,X30),k3_filter_1(X28,X29,X31))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_filter_1])])])])).
cnf(c_0_43, plain, (r1_binop_1(X1,X2,X3)|~r3_binop_1(X1,X2,X3)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_44, plain, ![X36, X37, X38]:((v1_funct_1(k3_filter_1(X36,X37,X38))|(v1_xboole_0(X36)|(~v1_partfun1(X37,X36)|(~v3_relat_2(X37)|(~v8_relat_2(X37)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(X36,X36)))))|(~v1_funct_1(X38)|(~v1_funct_2(X38,k2_zfmisc_1(X36,X36),X36)|~m1_subset_1(X38,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X36,X36),X36))))))))&((v1_funct_2(k3_filter_1(X36,X37,X38),k2_zfmisc_1(k8_eqrel_1(X36,X37),k8_eqrel_1(X36,X37)),k8_eqrel_1(X36,X37))|(v1_xboole_0(X36)|(~v1_partfun1(X37,X36)|(~v3_relat_2(X37)|(~v8_relat_2(X37)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(X36,X36)))))|(~v1_funct_1(X38)|(~v1_funct_2(X38,k2_zfmisc_1(X36,X36),X36)|~m1_subset_1(X38,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X36,X36),X36))))))))&(m1_subset_1(k3_filter_1(X36,X37,X38),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k8_eqrel_1(X36,X37),k8_eqrel_1(X36,X37)),k8_eqrel_1(X36,X37))))|(v1_xboole_0(X36)|(~v1_partfun1(X37,X36)|(~v3_relat_2(X37)|(~v8_relat_2(X37)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(X36,X36)))))|(~v1_funct_1(X38)|(~v1_funct_2(X38,k2_zfmisc_1(X36,X36),X36)|~m1_subset_1(X38,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X36,X36),X36)))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_filter_1])])])])).
fof(c_0_45, plain, ![X39, X40, X41]:((~m2_subset_1(X41,X39,X40)|m1_subset_1(X41,X40)|(v1_xboole_0(X39)|(v1_xboole_0(X40)|~m1_subset_1(X40,k1_zfmisc_1(X39)))))&(~m1_subset_1(X41,X40)|m2_subset_1(X41,X39,X40)|(v1_xboole_0(X39)|(v1_xboole_0(X40)|~m1_subset_1(X40,k1_zfmisc_1(X39)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])).
fof(c_0_46, plain, ![X42, X43]:(~v1_xboole_0(X42)|(~m1_subset_1(X43,k1_zfmisc_1(X42))|v1_xboole_0(X43))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
cnf(c_0_47, plain, (v1_xboole_0(X1)|m2_subset_1(k9_eqrel_1(X1,X2,X3),k1_zfmisc_1(X1),k8_eqrel_1(X1,X2))|~v3_relat_2(X2)|~v8_relat_2(X2)|~v1_partfun1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_48, negated_conjecture, (v1_partfun1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_49, negated_conjecture, (v8_relat_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_50, negated_conjecture, (v3_relat_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_51, plain, ![X44, X45]:(~m1_eqrel_1(X45,X44)|m1_subset_1(X45,k1_zfmisc_1(k1_zfmisc_1(X44)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_eqrel_1])])).
cnf(c_0_52, plain, (m1_eqrel_1(k8_eqrel_1(X2,X1),X2)|~v3_relat_2(X1)|~v8_relat_2(X1)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_53, plain, (v1_xboole_0(X1)|~v3_relat_2(X2)|~v8_relat_2(X2)|~v1_partfun1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_xboole_0(k7_eqrel_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_54, plain, (k8_eqrel_1(X2,X1)=k7_eqrel_1(X2,X1)|~v3_relat_2(X1)|~v8_relat_2(X1)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_55, plain, (v1_xboole_0(X1)|r2_binop_1(k8_eqrel_1(X1,X2),k9_eqrel_1(X1,X2,X3),k3_filter_1(X1,X2,X4))|~v1_partfun1(X2,X1)|~v3_relat_2(X2)|~v8_relat_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X3,X1)|~m2_filter_1(X4,X1,X2)|~r2_binop_1(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_56, negated_conjecture, (r2_binop_1(esk1_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_41])])).
cnf(c_0_57, plain, (v1_xboole_0(X1)|r1_binop_1(k8_eqrel_1(X1,X2),k9_eqrel_1(X1,X2,X3),k3_filter_1(X1,X2,X4))|~v1_partfun1(X2,X1)|~v3_relat_2(X2)|~v8_relat_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X3,X1)|~m2_filter_1(X4,X1,X2)|~r1_binop_1(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_58, negated_conjecture, (r1_binop_1(esk1_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_41])])).
cnf(c_0_59, plain, (v1_funct_1(k3_filter_1(X1,X2,X3))|v1_xboole_0(X1)|~v1_partfun1(X2,X1)|~v3_relat_2(X2)|~v8_relat_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_60, plain, (m1_subset_1(X1,X3)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_61, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_62, negated_conjecture, (m2_subset_1(k9_eqrel_1(esk1_0,esk2_0,X1),k1_zfmisc_1(esk1_0),k8_eqrel_1(esk1_0,esk2_0))|~m1_subset_1(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_21]), c_0_49]), c_0_50])]), c_0_25])).
cnf(c_0_63, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_eqrel_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_64, negated_conjecture, (m1_eqrel_1(k8_eqrel_1(esk1_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_48]), c_0_21]), c_0_49]), c_0_50])])).
cnf(c_0_65, negated_conjecture, (~v1_xboole_0(k7_eqrel_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_48]), c_0_21]), c_0_49]), c_0_50])]), c_0_25])).
cnf(c_0_66, negated_conjecture, (k7_eqrel_1(esk1_0,esk2_0)=k8_eqrel_1(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_48]), c_0_21]), c_0_49]), c_0_50])])).
cnf(c_0_67, negated_conjecture, (r2_binop_1(k8_eqrel_1(esk1_0,X1),k9_eqrel_1(esk1_0,X1,esk3_0),k3_filter_1(esk1_0,X1,esk4_0))|~m2_filter_1(esk4_0,esk1_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))|~v8_relat_2(X1)|~v3_relat_2(X1)|~v1_partfun1(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_41])]), c_0_25])).
cnf(c_0_68, negated_conjecture, (r1_binop_1(k8_eqrel_1(esk1_0,X1),k9_eqrel_1(esk1_0,X1,esk3_0),k3_filter_1(esk1_0,X1,esk4_0))|~m2_filter_1(esk4_0,esk1_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))|~v8_relat_2(X1)|~v3_relat_2(X1)|~v1_partfun1(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_41])]), c_0_25])).
cnf(c_0_69, negated_conjecture, (v1_funct_1(k3_filter_1(esk1_0,X1,esk4_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))|~v8_relat_2(X1)|~v3_relat_2(X1)|~v1_partfun1(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_38]), c_0_39]), c_0_40])]), c_0_25])).
cnf(c_0_70, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~m2_subset_1(X1,X3,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(csr,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_71, negated_conjecture, (m2_subset_1(k9_eqrel_1(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk1_0),k8_eqrel_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_62, c_0_41])).
cnf(c_0_72, negated_conjecture, (m1_subset_1(k8_eqrel_1(esk1_0,esk2_0),k1_zfmisc_1(k1_zfmisc_1(esk1_0)))), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_73, negated_conjecture, (~v1_xboole_0(k8_eqrel_1(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_74, plain, (v1_funct_2(k3_filter_1(X1,X2,X3),k2_zfmisc_1(k8_eqrel_1(X1,X2),k8_eqrel_1(X1,X2)),k8_eqrel_1(X1,X2))|v1_xboole_0(X1)|~v1_partfun1(X2,X1)|~v3_relat_2(X2)|~v8_relat_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_75, plain, (r3_binop_1(X1,X2,X3)|~r1_binop_1(X1,X2,X3)|~r2_binop_1(X1,X2,X3)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_76, negated_conjecture, (r2_binop_1(k8_eqrel_1(esk1_0,esk2_0),k9_eqrel_1(esk1_0,esk2_0,esk3_0),k3_filter_1(esk1_0,esk2_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_24]), c_0_21]), c_0_49]), c_0_50]), c_0_48])])).
cnf(c_0_77, negated_conjecture, (r1_binop_1(k8_eqrel_1(esk1_0,esk2_0),k9_eqrel_1(esk1_0,esk2_0,esk3_0),k3_filter_1(esk1_0,esk2_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_24]), c_0_21]), c_0_49]), c_0_50]), c_0_48])])).
cnf(c_0_78, negated_conjecture, (v1_funct_1(k3_filter_1(esk1_0,esk2_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_48]), c_0_21]), c_0_49]), c_0_50])])).
cnf(c_0_79, negated_conjecture, (m1_subset_1(k9_eqrel_1(esk1_0,esk2_0,esk3_0),k8_eqrel_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_72])]), c_0_73])).
cnf(c_0_80, negated_conjecture, (~r3_binop_1(k8_eqrel_1(esk1_0,esk2_0),k9_eqrel_1(esk1_0,esk2_0,esk3_0),k3_filter_1(esk1_0,esk2_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_81, negated_conjecture, (v1_funct_2(k3_filter_1(esk1_0,X1,esk4_0),k2_zfmisc_1(k8_eqrel_1(esk1_0,X1),k8_eqrel_1(esk1_0,X1)),k8_eqrel_1(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))|~v8_relat_2(X1)|~v3_relat_2(X1)|~v1_partfun1(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_38]), c_0_39]), c_0_40])]), c_0_25])).
cnf(c_0_82, plain, (m1_subset_1(k3_filter_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k8_eqrel_1(X1,X2),k8_eqrel_1(X1,X2)),k8_eqrel_1(X1,X2))))|v1_xboole_0(X1)|~v1_partfun1(X2,X1)|~v3_relat_2(X2)|~v8_relat_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_83, negated_conjecture, (~v1_funct_2(k3_filter_1(esk1_0,esk2_0,esk4_0),k2_zfmisc_1(k8_eqrel_1(esk1_0,esk2_0),k8_eqrel_1(esk1_0,esk2_0)),k8_eqrel_1(esk1_0,esk2_0))|~m1_subset_1(k3_filter_1(esk1_0,esk2_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k8_eqrel_1(esk1_0,esk2_0),k8_eqrel_1(esk1_0,esk2_0)),k8_eqrel_1(esk1_0,esk2_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77]), c_0_78]), c_0_79])]), c_0_80])).
cnf(c_0_84, negated_conjecture, (v1_funct_2(k3_filter_1(esk1_0,esk2_0,esk4_0),k2_zfmisc_1(k8_eqrel_1(esk1_0,esk2_0),k8_eqrel_1(esk1_0,esk2_0)),k8_eqrel_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_48]), c_0_21]), c_0_49]), c_0_50])])).
cnf(c_0_85, negated_conjecture, (m1_subset_1(k3_filter_1(esk1_0,X1,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k8_eqrel_1(esk1_0,X1),k8_eqrel_1(esk1_0,X1)),k8_eqrel_1(esk1_0,X1))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))|~v8_relat_2(X1)|~v3_relat_2(X1)|~v1_partfun1(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_38]), c_0_39]), c_0_40])]), c_0_25])).
cnf(c_0_86, negated_conjecture, (~m1_subset_1(k3_filter_1(esk1_0,esk2_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k8_eqrel_1(esk1_0,esk2_0),k8_eqrel_1(esk1_0,esk2_0)),k8_eqrel_1(esk1_0,esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_83, c_0_84])])).
cnf(c_0_87, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_48]), c_0_21]), c_0_49]), c_0_50])]), c_0_86]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 88
# Proof object clause steps            : 57
# Proof object formula steps           : 31
# Proof object conjectures             : 39
# Proof object clause conjectures      : 36
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 29
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 24
# Proof object simplifying inferences  : 92
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 99
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 98
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 41
# ...of the previous two non-trivial   : 42
# Contextual simplify-reflections      : 2
# Paramodulations                      : 41
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 65
#    Positive orientable unit clauses  : 24
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 36
# Current number of unprocessed clauses: 3
# ...number of literals in the above   : 18
# Current number of archived formulas  : 0
# Current number of archived clauses   : 33
# Clause-clause subsumption calls (NU) : 920
# Rec. Clause-clause subsumption calls : 68
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 26
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5293

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.003 s
# Total time               : 0.013 s
# Maximum resident set size: 28636 pages
