# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t64_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>((k1_relat_1(X1)=k1_xboole_0|k2_relat_1(X1)=k1_xboole_0)=>X1=k1_xboole_0)), file('data/MPTP/proofs/79188dcdcdfe47af9d63037b07abcaac.E_input', t64_relat_1)).
fof(s3_funct_1__e2_24__funct_1, axiom, ![X1, X2]:?[X3]:((v1_relat_1(X3)&v1_funct_1(X3))&(k1_relat_1(X3)=X1&![X4]:(r2_hidden(X4,X1)=>k1_funct_1(X3,X4)=X2))), file('data/MPTP/proofs/79188dcdcdfe47af9d63037b07abcaac.E_input', s3_funct_1__e2_24__funct_1)).
fof(d5_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4))))), file('data/MPTP/proofs/79188dcdcdfe47af9d63037b07abcaac.E_input', d5_funct_1)).
fof(t15_funct_1, conjecture, ![X1]:(~(X1=k1_xboole_0)=>![X2]:?[X3]:((v1_relat_1(X3)&v1_funct_1(X3))&(k1_relat_1(X3)=X1&k2_relat_1(X3)=k1_tarski(X2)))), file('data/MPTP/proofs/79188dcdcdfe47af9d63037b07abcaac.E_input', t15_funct_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/79188dcdcdfe47af9d63037b07abcaac.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/79188dcdcdfe47af9d63037b07abcaac.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/79188dcdcdfe47af9d63037b07abcaac.E_input', t71_enumset1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/79188dcdcdfe47af9d63037b07abcaac.E_input', d1_tarski)).
fof(c_0_8, plain, ![X5]:((k1_relat_1(X5)!=k1_xboole_0|X5=k1_xboole_0|~v1_relat_1(X5))&(k2_relat_1(X5)!=k1_xboole_0|X5=k1_xboole_0|~v1_relat_1(X5))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t64_relat_1])])])).
fof(c_0_9, plain, ![X6, X7, X9]:((v1_relat_1(esk1_2(X6,X7))&v1_funct_1(esk1_2(X6,X7)))&(k1_relat_1(esk1_2(X6,X7))=X6&(~r2_hidden(X9,X6)|k1_funct_1(esk1_2(X6,X7),X9)=X7))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[s3_funct_1__e2_24__funct_1])])])])).
fof(c_0_10, plain, ![X23, X24, X25, X27, X28, X29, X31]:((((r2_hidden(esk3_3(X23,X24,X25),k1_relat_1(X23))|~r2_hidden(X25,X24)|X24!=k2_relat_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23)))&(X25=k1_funct_1(X23,esk3_3(X23,X24,X25))|~r2_hidden(X25,X24)|X24!=k2_relat_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23))))&(~r2_hidden(X28,k1_relat_1(X23))|X27!=k1_funct_1(X23,X28)|r2_hidden(X27,X24)|X24!=k2_relat_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23))))&((~r2_hidden(esk4_2(X23,X29),X29)|(~r2_hidden(X31,k1_relat_1(X23))|esk4_2(X23,X29)!=k1_funct_1(X23,X31))|X29=k2_relat_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23)))&((r2_hidden(esk5_2(X23,X29),k1_relat_1(X23))|r2_hidden(esk4_2(X23,X29),X29)|X29=k2_relat_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23)))&(esk4_2(X23,X29)=k1_funct_1(X23,esk5_2(X23,X29))|r2_hidden(esk4_2(X23,X29),X29)|X29=k2_relat_1(X23)|(~v1_relat_1(X23)|~v1_funct_1(X23)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_funct_1])])])])])])).
cnf(c_0_11, plain, (X1=k1_xboole_0|k1_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (k1_relat_1(esk1_2(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (v1_relat_1(esk1_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (r2_hidden(X3,X4)|~r2_hidden(X1,k1_relat_1(X2))|X3!=k1_funct_1(X2,X1)|X4!=k2_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, plain, (k1_funct_1(esk1_2(X2,X3),X1)=X3|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, plain, (v1_funct_1(esk1_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, plain, (esk1_2(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])])).
cnf(c_0_18, plain, (r2_hidden(X1,X2)|X2!=k2_relat_1(esk1_2(X3,X1))|~r2_hidden(X4,X3)), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_12]), c_0_16]), c_0_13])])])).
cnf(c_0_19, plain, (k1_relat_1(k1_xboole_0)=X1|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_12, c_0_17])).
cnf(c_0_20, plain, (v1_funct_1(k1_xboole_0)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_21, plain, (v1_relat_1(k1_xboole_0)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_13, c_0_17])).
fof(c_0_22, negated_conjecture, ~(![X1]:(~(X1=k1_xboole_0)=>![X2]:?[X3]:((v1_relat_1(X3)&v1_funct_1(X3))&(k1_relat_1(X3)=X1&k2_relat_1(X3)=k1_tarski(X2))))), inference(assume_negation,[status(cth)],[t15_funct_1])).
cnf(c_0_23, plain, (r2_hidden(X1,X2)|X2!=k2_relat_1(k1_xboole_0)|X3!=k1_xboole_0|~r2_hidden(X4,X3)), inference(spm,[status(thm)],[c_0_18, c_0_17])).
cnf(c_0_24, plain, (r2_hidden(esk3_3(X1,X2,X3),k1_relat_1(X1))|~r2_hidden(X3,X2)|X2!=k2_relat_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_25, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (v1_funct_1(k1_xboole_0)), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (v1_relat_1(k1_xboole_0)), inference(er,[status(thm)],[c_0_21])).
fof(c_0_28, negated_conjecture, ![X35]:(esk6_0!=k1_xboole_0&(~v1_relat_1(X35)|~v1_funct_1(X35)|(k1_relat_1(X35)!=esk6_0|k2_relat_1(X35)!=k1_tarski(esk7_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_22])])])])])).
fof(c_0_29, plain, ![X17]:k2_tarski(X17,X17)=k1_tarski(X17), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_30, plain, ![X18, X19]:k1_enumset1(X18,X18,X19)=k2_tarski(X18,X19), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_31, plain, ![X20, X21, X22]:k2_enumset1(X20,X20,X21,X22)=k1_enumset1(X20,X21,X22), inference(variable_rename,[status(thm)],[t71_enumset1])).
cnf(c_0_32, plain, (r2_hidden(X1,k2_relat_1(k1_xboole_0))|X2!=k1_xboole_0|~r2_hidden(X3,X2)), inference(er,[status(thm)],[c_0_23])).
cnf(c_0_33, plain, (r2_hidden(esk3_3(k1_xboole_0,X1,X2),k1_xboole_0)|X1!=k2_relat_1(k1_xboole_0)|~r2_hidden(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26]), c_0_27])])).
cnf(c_0_34, negated_conjecture, (~v1_relat_1(X1)|~v1_funct_1(X1)|k1_relat_1(X1)!=esk6_0|k2_relat_1(X1)!=k1_tarski(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, plain, (r2_hidden(X1,k2_relat_1(k1_xboole_0))|X2!=k2_relat_1(k1_xboole_0)|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_39, negated_conjecture, (k1_relat_1(X1)!=esk6_0|k2_relat_1(X1)!=k2_enumset1(esk7_0,esk7_0,esk7_0,esk7_0)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_37])).
cnf(c_0_40, plain, (k1_funct_1(k1_xboole_0,X1)=X2|X3!=k1_xboole_0|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_15, c_0_17])).
cnf(c_0_41, plain, (r2_hidden(esk5_2(X1,X2),k1_relat_1(X1))|r2_hidden(esk4_2(X1,X2),X2)|X2=k2_relat_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_42, plain, (X1=k1_funct_1(X2,esk3_3(X2,X3,X1))|~r2_hidden(X1,X3)|X3!=k2_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_43, plain, (r2_hidden(X1,k2_relat_1(k1_xboole_0))|~r2_hidden(X2,k2_relat_1(k1_xboole_0))), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_44, plain, (r2_hidden(k1_funct_1(X1,X2),X3)|X3!=k2_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_14])).
cnf(c_0_45, negated_conjecture, (k2_relat_1(esk1_2(X1,X2))!=k2_enumset1(esk7_0,esk7_0,esk7_0,esk7_0)|X1!=esk6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_12]), c_0_16]), c_0_13])])).
cnf(c_0_46, plain, (k1_funct_1(k1_xboole_0,esk5_2(X1,X2))=X3|X2=k2_relat_1(X1)|r2_hidden(esk4_2(X1,X2),X2)|k1_relat_1(X1)!=k1_xboole_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_47, plain, (X1=X2|X3!=k2_relat_1(esk1_2(X4,X2))|~r2_hidden(esk3_3(esk1_2(X4,X2),X3,X1),X4)|~r2_hidden(X1,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_42]), c_0_16]), c_0_13])])).
cnf(c_0_48, plain, (r2_hidden(X1,k2_relat_1(k1_xboole_0))|k2_relat_1(k1_xboole_0)!=k2_relat_1(X2)|~r2_hidden(X3,k1_relat_1(X2))|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_49, negated_conjecture, (X1=k2_relat_1(X2)|r2_hidden(esk4_2(X2,X1),X1)|k1_relat_1(X2)!=k1_xboole_0|X3!=esk6_0|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_46])).
cnf(c_0_50, plain, (X1=X2|X3!=k2_relat_1(k1_xboole_0)|X4!=k1_xboole_0|~r2_hidden(esk3_3(k1_xboole_0,X3,X1),X4)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_17]), c_0_23])).
cnf(c_0_51, plain, (r2_hidden(X1,k2_relat_1(k1_xboole_0))|~r2_hidden(X2,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_25]), c_0_26]), c_0_27])])).
cnf(c_0_52, negated_conjecture, (X1=k2_relat_1(X2)|r2_hidden(esk4_2(X2,X1),X1)|k1_relat_1(X2)!=k1_xboole_0|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_49])).
cnf(c_0_53, plain, (X1=X2|X3!=k2_relat_1(k1_xboole_0)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_50, c_0_33])).
cnf(c_0_54, negated_conjecture, (k1_xboole_0=k2_relat_1(X1)|r2_hidden(X2,k2_relat_1(k1_xboole_0))|k1_relat_1(X1)!=k1_xboole_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_55, plain, (X1=X2|~r2_hidden(X1,k2_relat_1(k1_xboole_0))), inference(er,[status(thm)],[c_0_53])).
cnf(c_0_56, negated_conjecture, (k2_relat_1(k1_xboole_0)=k1_xboole_0|r2_hidden(X1,k2_relat_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_25]), c_0_26]), c_0_27])])).
cnf(c_0_57, negated_conjecture, (k2_relat_1(k1_xboole_0)=k1_xboole_0|X1=X2), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_58, negated_conjecture, (k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(csr,[status(thm)],[inference(ef,[status(thm)],[c_0_57]), c_0_57])).
cnf(c_0_59, negated_conjecture, (k1_xboole_0=k2_relat_1(X1)|r2_hidden(X2,k1_xboole_0)|k1_relat_1(X1)!=k1_xboole_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_54, c_0_58])).
cnf(c_0_60, negated_conjecture, (k2_relat_1(esk1_2(X1,X2))=k1_xboole_0|r2_hidden(X3,k1_xboole_0)|X1!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_12]), c_0_16]), c_0_13])])).
cnf(c_0_61, plain, (X1=X2|X3!=k1_xboole_0|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_53, c_0_58])).
cnf(c_0_62, negated_conjecture, (k2_relat_1(esk1_2(k1_xboole_0,X1))=k1_xboole_0|r2_hidden(X2,k1_xboole_0)), inference(er,[status(thm)],[c_0_60])).
cnf(c_0_63, negated_conjecture, (k2_relat_1(esk1_2(k1_xboole_0,X1))=k1_xboole_0|X2=X3), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_64, plain, (X1=k1_xboole_0|k2_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_65, negated_conjecture, (k2_relat_1(esk1_2(k1_xboole_0,X1))=k1_xboole_0), inference(csr,[status(thm)],[inference(ef,[status(thm)],[c_0_63]), c_0_63])).
cnf(c_0_66, plain, (r2_hidden(esk3_3(esk1_2(X1,X2),X3,X4),X1)|X3!=k2_relat_1(esk1_2(X1,X2))|~r2_hidden(X4,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_12]), c_0_16]), c_0_13])])).
fof(c_0_67, plain, ![X10, X11, X12, X13, X14, X15]:(((~r2_hidden(X12,X11)|X12=X10|X11!=k1_tarski(X10))&(X13!=X10|r2_hidden(X13,X11)|X11!=k1_tarski(X10)))&((~r2_hidden(esk2_2(X14,X15),X15)|esk2_2(X14,X15)!=X14|X15=k1_tarski(X14))&(r2_hidden(esk2_2(X14,X15),X15)|esk2_2(X14,X15)=X14|X15=k1_tarski(X14)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_68, negated_conjecture, (esk1_2(k1_xboole_0,X1)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_13])])).
cnf(c_0_69, plain, (X1=X2|X3!=k2_relat_1(esk1_2(X4,X2))|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_47, c_0_66])).
cnf(c_0_70, plain, (r2_hidden(esk2_2(X1,X2),X2)|esk2_2(X1,X2)=X1|X2=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_71, negated_conjecture, (k1_funct_1(k1_xboole_0,X1)=X2|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_15, c_0_68])).
cnf(c_0_72, plain, (X1=X2|~r2_hidden(X1,k2_relat_1(esk1_2(X3,X2)))), inference(er,[status(thm)],[c_0_69])).
cnf(c_0_73, plain, (esk2_2(X1,X2)=X1|X2=k2_enumset1(X1,X1,X1,X1)|r2_hidden(esk2_2(X1,X2),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_35]), c_0_36]), c_0_37])).
cnf(c_0_74, negated_conjecture, (X1=X2|~r2_hidden(X3,k1_xboole_0)), inference(spm,[status(thm)],[c_0_71, c_0_71])).
cnf(c_0_75, plain, (X1=k2_relat_1(esk1_2(X2,X3))|r2_hidden(esk4_2(esk1_2(X2,X3),X1),X1)|r2_hidden(esk5_2(esk1_2(X2,X3),X1),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_12]), c_0_16]), c_0_13])])).
cnf(c_0_76, plain, (X2=k1_tarski(X1)|~r2_hidden(esk2_2(X1,X2),X2)|esk2_2(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_77, plain, (k2_relat_1(esk1_2(X1,X2))=k2_enumset1(X3,X3,X3,X3)|esk2_2(X3,k2_relat_1(esk1_2(X1,X2)))=X3|esk2_2(X3,k2_relat_1(esk1_2(X1,X2)))=X2), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_78, negated_conjecture, (X1=k1_xboole_0|X2=X3|r2_hidden(esk4_2(k1_xboole_0,X1),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_65]), c_0_68])).
cnf(c_0_79, plain, (X2=k2_enumset1(X1,X1,X1,X1)|esk2_2(X1,X2)!=X1|~r2_hidden(esk2_2(X1,X2),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_35]), c_0_36]), c_0_37])).
cnf(c_0_80, plain, (k2_relat_1(esk1_2(X1,X2))=k2_enumset1(X2,X2,X2,X2)|esk2_2(X2,k2_relat_1(esk1_2(X1,X2)))=X2), inference(er,[status(thm)],[inference(ef,[status(thm)],[c_0_77])])).
cnf(c_0_81, plain, (r2_hidden(X1,k2_relat_1(esk1_2(X2,X1)))|~r2_hidden(X3,X2)), inference(er,[status(thm)],[c_0_18])).
cnf(c_0_82, negated_conjecture, (X1=k1_xboole_0|r2_hidden(esk4_2(k1_xboole_0,X1),X1)), inference(csr,[status(thm)],[inference(ef,[status(thm)],[c_0_78]), c_0_78])).
cnf(c_0_83, plain, (k2_relat_1(esk1_2(X1,X2))=k2_enumset1(X2,X2,X2,X2)|~r2_hidden(X2,k2_relat_1(esk1_2(X1,X2)))), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_84, negated_conjecture, (X1=k1_xboole_0|r2_hidden(X2,k2_relat_1(esk1_2(X1,X2)))), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_85, negated_conjecture, (k2_relat_1(esk1_2(X1,X2))=k2_enumset1(X2,X2,X2,X2)|X1=k1_xboole_0), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_86, negated_conjecture, (X1=k1_xboole_0|k2_enumset1(X2,X2,X2,X2)!=k2_enumset1(esk7_0,esk7_0,esk7_0,esk7_0)|X1!=esk6_0), inference(spm,[status(thm)],[c_0_45, c_0_85])).
cnf(c_0_87, negated_conjecture, (X1=k1_xboole_0|X1!=esk6_0), inference(er,[status(thm)],[c_0_86])).
cnf(c_0_88, negated_conjecture, (esk6_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_89, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_87]), c_0_88]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 90
# Proof object clause steps            : 73
# Proof object formula steps           : 17
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 51
# Proof object simplifying inferences  : 53
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 18
# Processed clauses                    : 1631
# ...of these trivial                  : 14
# ...subsumed                          : 1044
# ...remaining for further processing  : 573
# Other redundant clauses eliminated   : 468
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 42
# Backward-rewritten                   : 48
# Generated clauses                    : 7125
# ...of the previous two non-trivial   : 6341
# Contextual simplify-reflections      : 16
# Paramodulations                      : 6366
# Factorizations                       : 40
# Equation resolutions                 : 719
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 464
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 450
# Current number of unprocessed clauses: 3157
# ...number of literals in the above   : 19739
# Current number of archived formulas  : 0
# Current number of archived clauses   : 111
# Clause-clause subsumption calls (NU) : 72514
# Rec. Clause-clause subsumption calls : 13137
# Non-unit clause-clause subsumptions  : 1093
# Unit Clause-clause subsumption calls : 107
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 49
# BW rewrite match successes           : 8
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 92525

# -------------------------------------------------
# User time                : 0.118 s
# System time              : 0.003 s
# Total time               : 0.121 s
# Maximum resident set size: 28740 pages
