# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/7919fb83d64a4e4880d5d4460b17c219.E_input', t7_xboole_1)).
fof(idempotence_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/7919fb83d64a4e4880d5d4460b17c219.E_input', idempotence_k2_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/7919fb83d64a4e4880d5d4460b17c219.E_input', t28_xboole_1)).
fof(t134_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4)))), file('data/MPTP/proofs/7919fb83d64a4e4880d5d4460b17c219.E_input', t134_zfmisc_1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/7919fb83d64a4e4880d5d4460b17c219.E_input', t4_xboole_0)).
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/7919fb83d64a4e4880d5d4460b17c219.E_input', d2_zfmisc_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/7919fb83d64a4e4880d5d4460b17c219.E_input', commutativity_k3_xboole_0)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/7919fb83d64a4e4880d5d4460b17c219.E_input', commutativity_k2_xboole_0)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/7919fb83d64a4e4880d5d4460b17c219.E_input', t17_xboole_1)).
fof(t120_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&k2_zfmisc_1(X3,k2_xboole_0(X1,X2))=k2_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))), file('data/MPTP/proofs/7919fb83d64a4e4880d5d4460b17c219.E_input', t120_zfmisc_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/7919fb83d64a4e4880d5d4460b17c219.E_input', t12_xboole_1)).
fof(t123_zfmisc_1, axiom, ![X1, X2, X3, X4]:k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)), file('data/MPTP/proofs/7919fb83d64a4e4880d5d4460b17c219.E_input', t123_zfmisc_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/7919fb83d64a4e4880d5d4460b17c219.E_input', d7_xboole_0)).
fof(t117_zfmisc_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&((r1_tarski(k2_zfmisc_1(X2,X1),k2_zfmisc_1(X3,X1))|r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))&~(r1_tarski(X2,X3))))), file('data/MPTP/proofs/7919fb83d64a4e4880d5d4460b17c219.E_input', t117_zfmisc_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/7919fb83d64a4e4880d5d4460b17c219.E_input', t7_xboole_0)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/7919fb83d64a4e4880d5d4460b17c219.E_input', d10_xboole_0)).
fof(c_0_16, plain, ![X52, X53]:r1_tarski(X52,k2_xboole_0(X52,X53)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_17, plain, ![X54]:k2_xboole_0(X54,X54)=X54, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k2_xboole_0])])).
fof(c_0_18, plain, ![X29, X30]:(~r1_tarski(X29,X30)|k3_xboole_0(X29,X30)=X29), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_19, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, (k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_21, negated_conjecture, ~(![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4))))), inference(assume_negation,[status(cth)],[t134_zfmisc_1])).
fof(c_0_22, plain, ![X16, X17, X19, X20, X21]:((r1_xboole_0(X16,X17)|r2_hidden(esk5_2(X16,X17),k3_xboole_0(X16,X17)))&(~r2_hidden(X21,k3_xboole_0(X19,X20))|~r1_xboole_0(X19,X20))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
cnf(c_0_23, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
fof(c_0_25, plain, ![X31, X32, X33, X34, X37, X38, X39, X40, X41, X42, X44, X45]:((((r2_hidden(esk6_4(X31,X32,X33,X34),X31)|~r2_hidden(X34,X33)|X33!=k2_zfmisc_1(X31,X32))&((r2_hidden(esk7_4(X31,X32,X33,X34),X32)|~r2_hidden(X34,X33)|X33!=k2_zfmisc_1(X31,X32))&(X34=k4_tarski(esk6_4(X31,X32,X33,X34),esk7_4(X31,X32,X33,X34))|~r2_hidden(X34,X33)|X33!=k2_zfmisc_1(X31,X32))))&(~r2_hidden(X38,X31)|(~r2_hidden(X39,X32)|X37!=k4_tarski(X38,X39))|r2_hidden(X37,X33)|X33!=k2_zfmisc_1(X31,X32)))&((~r2_hidden(esk8_3(X40,X41,X42),X42)|(~r2_hidden(X44,X40)|(~r2_hidden(X45,X41)|esk8_3(X40,X41,X42)!=k4_tarski(X44,X45)))|X42=k2_zfmisc_1(X40,X41))&((r2_hidden(esk9_3(X40,X41,X42),X40)|r2_hidden(esk8_3(X40,X41,X42),X42)|X42=k2_zfmisc_1(X40,X41))&((r2_hidden(esk10_3(X40,X41,X42),X41)|r2_hidden(esk8_3(X40,X41,X42),X42)|X42=k2_zfmisc_1(X40,X41))&(esk8_3(X40,X41,X42)=k4_tarski(esk9_3(X40,X41,X42),esk10_3(X40,X41,X42))|r2_hidden(esk8_3(X40,X41,X42),X42)|X42=k2_zfmisc_1(X40,X41)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
fof(c_0_26, plain, ![X57, X58]:k3_xboole_0(X57,X58)=k3_xboole_0(X58,X57), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_27, plain, ![X50, X51]:k2_xboole_0(X50,X51)=k2_xboole_0(X51,X50), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_28, plain, ![X22, X23]:r1_tarski(k3_xboole_0(X22,X23),X22), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_29, plain, ![X13, X14, X15]:(k2_zfmisc_1(k2_xboole_0(X13,X14),X15)=k2_xboole_0(k2_zfmisc_1(X13,X15),k2_zfmisc_1(X14,X15))&k2_zfmisc_1(X15,k2_xboole_0(X13,X14))=k2_xboole_0(k2_zfmisc_1(X15,X13),k2_zfmisc_1(X15,X14))), inference(variable_rename,[status(thm)],[t120_zfmisc_1])).
fof(c_0_30, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k2_zfmisc_1(esk3_0,esk4_0)&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk1_0!=esk3_0|esk2_0!=esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).
cnf(c_0_31, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32, plain, (k3_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_33, plain, (r2_hidden(esk6_4(X1,X2,X3,X4),X1)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, plain, (r2_hidden(X5,X6)|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)|X5!=k4_tarski(X1,X3)|X6!=k2_zfmisc_1(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk5_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_36, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_37, plain, (r2_hidden(esk7_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_38, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_39, plain, ![X63, X64]:(~r1_tarski(X63,X64)|k2_xboole_0(X63,X64)=X64), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_40, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_41, plain, (k2_zfmisc_1(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_42, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k2_zfmisc_1(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_43, plain, ![X59, X60, X61, X62]:k2_zfmisc_1(k3_xboole_0(X59,X60),k3_xboole_0(X61,X62))=k3_xboole_0(k2_zfmisc_1(X59,X61),k2_zfmisc_1(X60,X62)), inference(variable_rename,[status(thm)],[t123_zfmisc_1])).
cnf(c_0_44, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_45, plain, (r2_hidden(esk6_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X1)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_33])).
cnf(c_0_46, plain, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_34])])).
cnf(c_0_47, plain, (r2_hidden(esk5_2(X1,X2),k3_xboole_0(X2,X1))|r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
fof(c_0_48, plain, ![X48, X49]:((~r1_xboole_0(X48,X49)|k3_xboole_0(X48,X49)=k1_xboole_0)&(k3_xboole_0(X48,X49)!=k1_xboole_0|r1_xboole_0(X48,X49))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_49, plain, (r2_hidden(esk7_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X2)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_50, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_19, c_0_38])).
cnf(c_0_51, plain, (k2_zfmisc_1(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_52, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_53, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_40, c_0_36])).
cnf(c_0_54, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(spm,[status(thm)],[c_0_23, c_0_19])).
cnf(c_0_55, negated_conjecture, (k2_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,X1))=k2_zfmisc_1(esk1_0,k2_xboole_0(X1,esk2_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_38])).
cnf(c_0_56, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_57, plain, (~r2_hidden(X1,k2_zfmisc_1(X2,X3))|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_58, negated_conjecture, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(esk3_0,esk4_0))|~r2_hidden(X2,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_46, c_0_42])).
cnf(c_0_59, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_31, c_0_47])).
cnf(c_0_60, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_61, plain, (~r2_hidden(X1,k2_zfmisc_1(X2,X3))|~r1_xboole_0(X3,X3)), inference(spm,[status(thm)],[c_0_44, c_0_49])).
fof(c_0_62, plain, ![X26, X27, X28]:((~r1_tarski(k2_zfmisc_1(X27,X26),k2_zfmisc_1(X28,X26))|r1_tarski(X27,X28)|X26=k1_xboole_0)&(~r1_tarski(k2_zfmisc_1(X26,X27),k2_zfmisc_1(X26,X28))|r1_tarski(X27,X28)|X26=k1_xboole_0)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t117_zfmisc_1])])])])).
cnf(c_0_63, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(k2_xboole_0(X3,X1),X2))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_64, plain, (k2_xboole_0(X1,k3_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_38])).
cnf(c_0_65, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,k2_xboole_0(X1,esk2_0)))=k2_zfmisc_1(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_66, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,k2_xboole_0(X2,X4)))=k2_zfmisc_1(k3_xboole_0(X1,X3),X2)), inference(spm,[status(thm)],[c_0_56, c_0_54])).
cnf(c_0_67, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_53]), c_0_36])).
cnf(c_0_68, negated_conjecture, (~r2_hidden(X1,esk2_0)|~r2_hidden(X2,esk1_0)|~r1_xboole_0(esk3_0,esk3_0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_69, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X2,X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_59, c_0_60])).
fof(c_0_70, plain, ![X55]:(X55=k1_xboole_0|r2_hidden(esk11_1(X55),X55)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_71, negated_conjecture, (~r2_hidden(X1,esk2_0)|~r2_hidden(X2,esk1_0)|~r1_xboole_0(esk4_0,esk4_0)), inference(spm,[status(thm)],[c_0_61, c_0_58])).
cnf(c_0_72, plain, (r1_tarski(X2,X3)|X1=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_73, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_74, plain, (r1_tarski(k2_zfmisc_1(k3_xboole_0(X1,X2),X3),k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_75, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk3_0,esk1_0),esk4_0)=k2_zfmisc_1(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_76, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,k3_xboole_0(X4,X2)))=k3_xboole_0(k2_zfmisc_1(X1,X4),k2_zfmisc_1(X3,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_67]), c_0_56])).
cnf(c_0_77, plain, (k3_xboole_0(X1,k2_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_23, c_0_50])).
cnf(c_0_78, negated_conjecture, (k2_xboole_0(k2_zfmisc_1(X1,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))=k2_zfmisc_1(k2_xboole_0(X1,esk1_0),esk2_0)), inference(spm,[status(thm)],[c_0_51, c_0_42])).
cnf(c_0_79, negated_conjecture, (k1_xboole_0!=esk3_0|~r2_hidden(X1,esk2_0)|~r2_hidden(X2,esk1_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_32])).
cnf(c_0_80, plain, (X1=k1_xboole_0|r2_hidden(esk11_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_81, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_82, negated_conjecture, (k1_xboole_0!=esk4_0|~r2_hidden(X1,esk2_0)|~r2_hidden(X2,esk1_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_69]), c_0_32])).
fof(c_0_83, plain, ![X11, X12]:(((r1_tarski(X11,X12)|X11!=X12)&(r1_tarski(X12,X11)|X11!=X12))&(~r1_tarski(X11,X12)|~r1_tarski(X12,X11)|X11=X12)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_84, negated_conjecture, (r1_tarski(esk2_0,X1)|~r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_42]), c_0_73])).
cnf(c_0_85, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,esk4_0))), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_86, plain, (r1_tarski(k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4)),k2_zfmisc_1(X1,X4))), inference(spm,[status(thm)],[c_0_40, c_0_76])).
cnf(c_0_87, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(k2_xboole_0(X1,esk1_0),esk2_0))=k2_zfmisc_1(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_88, negated_conjecture, (k1_xboole_0!=esk3_0|~r2_hidden(X1,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_81])).
cnf(c_0_89, negated_conjecture, (k1_xboole_0!=esk4_0|~r2_hidden(X1,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_80]), c_0_81])).
cnf(c_0_90, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_91, negated_conjecture, (r1_tarski(esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_92, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk3_0,esk2_0))), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_93, negated_conjecture, (k1_xboole_0!=esk3_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_80]), c_0_73])).
cnf(c_0_94, plain, (r1_tarski(X1,X3)|X2=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_95, negated_conjecture, (k1_xboole_0!=esk4_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_80]), c_0_73])).
cnf(c_0_96, negated_conjecture, (esk2_0=esk4_0|~r1_tarski(esk4_0,esk2_0)), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_97, negated_conjecture, (r1_tarski(esk4_0,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_92]), c_0_93])).
cnf(c_0_98, negated_conjecture, (r1_tarski(esk3_0,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_85]), c_0_95])).
cnf(c_0_99, negated_conjecture, (r1_tarski(esk1_0,X1)|~r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(X1,esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_42]), c_0_81])).
cnf(c_0_100, negated_conjecture, (esk1_0!=esk3_0|esk2_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_101, negated_conjecture, (esk2_0=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_96, c_0_97])])).
cnf(c_0_102, negated_conjecture, (esk1_0=esk3_0|~r1_tarski(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_90, c_0_98])).
cnf(c_0_103, negated_conjecture, (r1_tarski(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_99, c_0_92])).
cnf(c_0_104, negated_conjecture, (esk1_0!=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_100, c_0_101])])).
cnf(c_0_105, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_102, c_0_103])]), c_0_104]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 106
# Proof object clause steps            : 73
# Proof object formula steps           : 33
# Proof object conjectures             : 34
# Proof object clause conjectures      : 31
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 43
# Proof object simplifying inferences  : 25
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 35
# Processed clauses                    : 3029
# ...of these trivial                  : 218
# ...subsumed                          : 2261
# ...remaining for further processing  : 549
# Other redundant clauses eliminated   : 47
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 24
# Backward-rewritten                   : 196
# Generated clauses                    : 15632
# ...of the previous two non-trivial   : 11642
# Contextual simplify-reflections      : 1
# Paramodulations                      : 15568
# Factorizations                       : 18
# Equation resolutions                 : 47
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 317
#    Positive orientable unit clauses  : 110
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 10
#    Non-unit-clauses                  : 195
# Current number of unprocessed clauses: 8395
# ...number of literals in the above   : 19641
# Current number of archived formulas  : 0
# Current number of archived clauses   : 220
# Clause-clause subsumption calls (NU) : 19880
# Rec. Clause-clause subsumption calls : 13266
# Non-unit clause-clause subsumptions  : 1674
# Unit Clause-clause subsumption calls : 2055
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 262
# BW rewrite match successes           : 29
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 162149

# -------------------------------------------------
# User time                : 0.109 s
# System time              : 0.007 s
# Total time               : 0.116 s
# Maximum resident set size: 28648 pages
