# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t16_funct_2, conjecture, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(![X4]:~((r2_hidden(X4,X2)&![X5]:~((r2_hidden(X5,X1)&X4=k1_funct_1(X3,X5)))))=>k2_relset_1(X1,X2,X3)=X2)), file('data/MPTP/proofs/792be7d4445c43c1a3805469d96015b0.E_input', t16_funct_2)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/792be7d4445c43c1a3805469d96015b0.E_input', t5_subset)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/792be7d4445c43c1a3805469d96015b0.E_input', t3_subset)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/792be7d4445c43c1a3805469d96015b0.E_input', d10_xboole_0)).
fof(d5_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4))))), file('data/MPTP/proofs/792be7d4445c43c1a3805469d96015b0.E_input', d5_funct_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/792be7d4445c43c1a3805469d96015b0.E_input', cc1_relset_1)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(((X2=k1_xboole_0=>X1=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X2,X3)))&(X2=k1_xboole_0=>(X1=k1_xboole_0|(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0))))), file('data/MPTP/proofs/792be7d4445c43c1a3805469d96015b0.E_input', d1_funct_2)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/792be7d4445c43c1a3805469d96015b0.E_input', redefinition_k1_relset_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/792be7d4445c43c1a3805469d96015b0.E_input', d3_tarski)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/792be7d4445c43c1a3805469d96015b0.E_input', redefinition_k2_relset_1)).
fof(d19_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(v5_relat_1(X2,X1)<=>r1_tarski(k2_relat_1(X2),X1))), file('data/MPTP/proofs/792be7d4445c43c1a3805469d96015b0.E_input', d19_relat_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('data/MPTP/proofs/792be7d4445c43c1a3805469d96015b0.E_input', cc2_relset_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/792be7d4445c43c1a3805469d96015b0.E_input', fc1_xboole_0)).
fof(c_0_13, negated_conjecture, ~(![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(![X4]:~((r2_hidden(X4,X2)&![X5]:~((r2_hidden(X5,X1)&X4=k1_funct_1(X3,X5)))))=>k2_relset_1(X1,X2,X3)=X2))), inference(assume_negation,[status(cth)],[t16_funct_2])).
fof(c_0_14, plain, ![X29, X30, X31]:(~r2_hidden(X29,X30)|(~m1_subset_1(X30,k1_zfmisc_1(X31))|~v1_xboole_0(X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_15, plain, ![X49, X50]:((~m1_subset_1(X49,k1_zfmisc_1(X50))|r1_tarski(X49,X50))&(~r1_tarski(X49,X50)|m1_subset_1(X49,k1_zfmisc_1(X50)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_16, plain, ![X41, X42]:(((r1_tarski(X41,X42)|X41!=X42)&(r1_tarski(X42,X41)|X41!=X42))&(~r1_tarski(X41,X42)|~r1_tarski(X42,X41)|X41=X42)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_17, plain, ![X14, X15, X16, X18, X19, X20, X22]:((((r2_hidden(esk5_3(X14,X15,X16),k1_relat_1(X14))|~r2_hidden(X16,X15)|X15!=k2_relat_1(X14)|(~v1_relat_1(X14)|~v1_funct_1(X14)))&(X16=k1_funct_1(X14,esk5_3(X14,X15,X16))|~r2_hidden(X16,X15)|X15!=k2_relat_1(X14)|(~v1_relat_1(X14)|~v1_funct_1(X14))))&(~r2_hidden(X19,k1_relat_1(X14))|X18!=k1_funct_1(X14,X19)|r2_hidden(X18,X15)|X15!=k2_relat_1(X14)|(~v1_relat_1(X14)|~v1_funct_1(X14))))&((~r2_hidden(esk6_2(X14,X20),X20)|(~r2_hidden(X22,k1_relat_1(X14))|esk6_2(X14,X20)!=k1_funct_1(X14,X22))|X20=k2_relat_1(X14)|(~v1_relat_1(X14)|~v1_funct_1(X14)))&((r2_hidden(esk7_2(X14,X20),k1_relat_1(X14))|r2_hidden(esk6_2(X14,X20),X20)|X20=k2_relat_1(X14)|(~v1_relat_1(X14)|~v1_funct_1(X14)))&(esk6_2(X14,X20)=k1_funct_1(X14,esk7_2(X14,X20))|r2_hidden(esk6_2(X14,X20),X20)|X20=k2_relat_1(X14)|(~v1_relat_1(X14)|~v1_funct_1(X14)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_funct_1])])])])])])).
fof(c_0_18, plain, ![X38, X39, X40]:(~m1_subset_1(X40,k1_zfmisc_1(k2_zfmisc_1(X38,X39)))|v1_relat_1(X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_19, negated_conjecture, ![X9]:((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,esk1_0,esk2_0)&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))&(((r2_hidden(esk4_1(X9),esk1_0)|~r2_hidden(X9,esk2_0))&(X9=k1_funct_1(esk3_0,esk4_1(X9))|~r2_hidden(X9,esk2_0)))&k2_relset_1(esk1_0,esk2_0,esk3_0)!=esk2_0)), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])])).
fof(c_0_20, plain, ![X11, X12, X13]:((((~v1_funct_2(X13,X11,X12)|X11=k1_relset_1(X11,X12,X13)|X12=k1_xboole_0|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12))))&(X11!=k1_relset_1(X11,X12,X13)|v1_funct_2(X13,X11,X12)|X12=k1_xboole_0|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12)))))&((~v1_funct_2(X13,X11,X12)|X11=k1_relset_1(X11,X12,X13)|X11!=k1_xboole_0|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12))))&(X11!=k1_relset_1(X11,X12,X13)|v1_funct_2(X13,X11,X12)|X11!=k1_xboole_0|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12))))))&((~v1_funct_2(X13,X11,X12)|X13=k1_xboole_0|X11=k1_xboole_0|X12!=k1_xboole_0|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12))))&(X13!=k1_xboole_0|v1_funct_2(X13,X11,X12)|X11=k1_xboole_0|X12!=k1_xboole_0|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
fof(c_0_21, plain, ![X35, X36, X37]:(~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(X35,X36)))|k1_relset_1(X35,X36,X37)=k1_relat_1(X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_22, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, plain, (r2_hidden(X3,X4)|~r2_hidden(X1,k1_relat_1(X2))|X3!=k1_funct_1(X2,X1)|X4!=k2_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, plain, (X2=k1_relset_1(X2,X3,X1)|X3=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, negated_conjecture, (v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_30, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, plain, (~v1_xboole_0(X1)|~r1_tarski(X2,X1)|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_32, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_24])).
fof(c_0_33, plain, ![X43, X44, X45, X46, X47]:((~r1_tarski(X43,X44)|(~r2_hidden(X45,X43)|r2_hidden(X45,X44)))&((r2_hidden(esk8_2(X46,X47),X46)|r1_tarski(X46,X47))&(~r2_hidden(esk8_2(X46,X47),X47)|r1_tarski(X46,X47)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_34, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relat_1(X1))|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))|~v1_funct_1(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_25])])).
cnf(c_0_35, negated_conjecture, (X1=k1_funct_1(esk3_0,esk4_1(X1))|~r2_hidden(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_36, negated_conjecture, (v1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_37, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_38, negated_conjecture, (k1_relset_1(esk1_0,esk2_0,esk3_0)=esk1_0|esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_27]), c_0_29])])).
cnf(c_0_39, negated_conjecture, (k1_relset_1(esk1_0,esk2_0,esk3_0)=k1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_30, c_0_27])).
cnf(c_0_40, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_41, plain, (r2_hidden(esk8_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_42, plain, ![X26, X27, X28]:(~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(X26,X27)))|k2_relset_1(X26,X27,X28)=k2_relat_1(X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
cnf(c_0_43, negated_conjecture, (r2_hidden(X1,k2_relat_1(esk3_0))|~r2_hidden(esk4_1(X1),k1_relat_1(esk3_0))|~r2_hidden(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_37])])).
cnf(c_0_44, negated_conjecture, (k1_relat_1(esk3_0)=esk1_0|esk2_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_45, negated_conjecture, (r2_hidden(esk4_1(X1),esk1_0)|~r2_hidden(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_46, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_47, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
fof(c_0_48, plain, ![X24, X25]:((~v5_relat_1(X25,X24)|r1_tarski(k2_relat_1(X25),X24)|~v1_relat_1(X25))&(~r1_tarski(k2_relat_1(X25),X24)|v5_relat_1(X25,X24)|~v1_relat_1(X25))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d19_relat_1])])])).
fof(c_0_49, plain, ![X32, X33, X34]:((v4_relat_1(X34,X32)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))&(v5_relat_1(X34,X33)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_50, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_51, plain, (r1_tarski(X1,X2)|~r2_hidden(esk8_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_52, negated_conjecture, (esk2_0=k1_xboole_0|r2_hidden(X1,k2_relat_1(esk3_0))|~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])).
cnf(c_0_53, plain, (X1=X2|~v1_xboole_0(X2)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_54, plain, (r1_tarski(k2_relat_1(X1),X2)|~v5_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_55, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_56, negated_conjecture, (k2_relset_1(esk1_0,esk2_0,esk3_0)!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_57, negated_conjecture, (k2_relset_1(esk1_0,esk2_0,esk3_0)=k2_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_50, c_0_27])).
cnf(c_0_58, negated_conjecture, (esk2_0=k1_xboole_0|r1_tarski(X1,k2_relat_1(esk3_0))|~r2_hidden(esk8_2(X1,k2_relat_1(esk3_0)),esk2_0)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_59, plain, (k2_relat_1(X1)=X2|~v1_xboole_0(X2)|~v5_relat_1(X1,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_60, negated_conjecture, (v5_relat_1(esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_55, c_0_27])).
cnf(c_0_61, negated_conjecture, (k2_relat_1(esk3_0)!=esk2_0), inference(rw,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_62, plain, (X1=k2_relat_1(X2)|~r1_tarski(X1,k2_relat_1(X2))|~v5_relat_1(X2,X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_46, c_0_54])).
cnf(c_0_63, negated_conjecture, (esk2_0=k1_xboole_0|r1_tarski(esk2_0,k2_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_58, c_0_41])).
cnf(c_0_64, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_36])]), c_0_61])).
cnf(c_0_65, negated_conjecture, (esk2_0=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_60]), c_0_36])]), c_0_61])).
cnf(c_0_66, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_67, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_65]), c_0_66])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 68
# Proof object clause steps            : 42
# Proof object formula steps           : 26
# Proof object conjectures             : 23
# Proof object clause conjectures      : 20
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 20
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 35
# Processed clauses                    : 236
# ...of these trivial                  : 4
# ...subsumed                          : 81
# ...remaining for further processing  : 151
# Other redundant clauses eliminated   : 11
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 7
# Backward-rewritten                   : 46
# Generated clauses                    : 330
# ...of the previous two non-trivial   : 314
# Contextual simplify-reflections      : 6
# Paramodulations                      : 321
# Factorizations                       : 0
# Equation resolutions                 : 11
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 89
#    Positive orientable unit clauses  : 18
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 69
# Current number of unprocessed clauses: 103
# ...number of literals in the above   : 403
# Current number of archived formulas  : 0
# Current number of archived clauses   : 53
# Clause-clause subsumption calls (NU) : 2607
# Rec. Clause-clause subsumption calls : 1648
# Non-unit clause-clause subsumptions  : 56
# Unit Clause-clause subsumption calls : 180
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 18
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6233

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.009 s
# Total time               : 0.015 s
# Maximum resident set size: 28676 pages
