# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t79_mcart_1, conjecture, ![X1, X2, X3, X4, X5, X6]:(m1_subset_1(X6,k4_zfmisc_1(X1,X2,X3,X4))=>(![X7]:(m1_subset_1(X7,X1)=>![X8]:(m1_subset_1(X8,X2)=>![X9]:(m1_subset_1(X9,X3)=>![X10]:(m1_subset_1(X10,X4)=>(X6=k4_mcart_1(X7,X8,X9,X10)=>X5=X7)))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|(X4=k1_xboole_0|X5=k8_mcart_1(X1,X2,X3,X4,X6))))))), file('data/MPTP/proofs/794bbfa394fe414a968885e21691b9fe.E_input', t79_mcart_1)).
fof(d4_mcart_1, axiom, ![X1, X2, X3, X4]:k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4), file('data/MPTP/proofs/794bbfa394fe414a968885e21691b9fe.E_input', d4_mcart_1)).
fof(l68_mcart_1, axiom, ![X1, X2, X3, X4]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&?[X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))&![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>![X9]:(m1_subset_1(X9,X4)=>~(X5=k4_mcart_1(X6,X7,X8,X9)))))))))))), file('data/MPTP/proofs/794bbfa394fe414a968885e21691b9fe.E_input', l68_mcart_1)).
fof(d4_zfmisc_1, axiom, ![X1, X2, X3, X4]:k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4), file('data/MPTP/proofs/794bbfa394fe414a968885e21691b9fe.E_input', d4_zfmisc_1)).
fof(t78_mcart_1, axiom, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&?[X6, X7, X8, X9]:(X5=k4_mcart_1(X6,X7,X8,X9)&~((k8_mcart_1(X1,X2,X3,X4,X5)=X6&(k9_mcart_1(X1,X2,X3,X4,X5)=X7&(k10_mcart_1(X1,X2,X3,X4,X5)=X8&k11_mcart_1(X1,X2,X3,X4,X5)=X9))))))))))), file('data/MPTP/proofs/794bbfa394fe414a968885e21691b9fe.E_input', t78_mcart_1)).
fof(c_0_5, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6]:(m1_subset_1(X6,k4_zfmisc_1(X1,X2,X3,X4))=>(![X7]:(m1_subset_1(X7,X1)=>![X8]:(m1_subset_1(X8,X2)=>![X9]:(m1_subset_1(X9,X3)=>![X10]:(m1_subset_1(X10,X4)=>(X6=k4_mcart_1(X7,X8,X9,X10)=>X5=X7)))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|(X4=k1_xboole_0|X5=k8_mcart_1(X1,X2,X3,X4,X6)))))))), inference(assume_negation,[status(cth)],[t79_mcart_1])).
fof(c_0_6, negated_conjecture, ![X17, X18, X19, X20]:(m1_subset_1(esk6_0,k4_zfmisc_1(esk1_0,esk2_0,esk3_0,esk4_0))&((~m1_subset_1(X17,esk1_0)|(~m1_subset_1(X18,esk2_0)|(~m1_subset_1(X19,esk3_0)|(~m1_subset_1(X20,esk4_0)|(esk6_0!=k4_mcart_1(X17,X18,X19,X20)|esk5_0=X17)))))&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&(esk4_0!=k1_xboole_0&esk5_0!=k8_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
fof(c_0_7, plain, ![X21, X22, X23, X24]:k4_mcart_1(X21,X22,X23,X24)=k4_tarski(k3_mcart_1(X21,X22,X23),X24), inference(variable_rename,[status(thm)],[d4_mcart_1])).
fof(c_0_8, plain, ![X25, X26, X27, X28, X29]:((m1_subset_1(esk7_5(X25,X26,X27,X28,X29),X25)|~m1_subset_1(X29,k4_zfmisc_1(X25,X26,X27,X28))|X28=k1_xboole_0|X27=k1_xboole_0|X26=k1_xboole_0|X25=k1_xboole_0)&((m1_subset_1(esk8_5(X25,X26,X27,X28,X29),X26)|~m1_subset_1(X29,k4_zfmisc_1(X25,X26,X27,X28))|X28=k1_xboole_0|X27=k1_xboole_0|X26=k1_xboole_0|X25=k1_xboole_0)&((m1_subset_1(esk9_5(X25,X26,X27,X28,X29),X27)|~m1_subset_1(X29,k4_zfmisc_1(X25,X26,X27,X28))|X28=k1_xboole_0|X27=k1_xboole_0|X26=k1_xboole_0|X25=k1_xboole_0)&((m1_subset_1(esk10_5(X25,X26,X27,X28,X29),X28)|~m1_subset_1(X29,k4_zfmisc_1(X25,X26,X27,X28))|X28=k1_xboole_0|X27=k1_xboole_0|X26=k1_xboole_0|X25=k1_xboole_0)&(X29=k4_mcart_1(esk7_5(X25,X26,X27,X28,X29),esk8_5(X25,X26,X27,X28,X29),esk9_5(X25,X26,X27,X28,X29),esk10_5(X25,X26,X27,X28,X29))|~m1_subset_1(X29,k4_zfmisc_1(X25,X26,X27,X28))|X28=k1_xboole_0|X27=k1_xboole_0|X26=k1_xboole_0|X25=k1_xboole_0))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l68_mcart_1])])])])])])])).
fof(c_0_9, plain, ![X34, X35, X36, X37]:k4_zfmisc_1(X34,X35,X36,X37)=k2_zfmisc_1(k3_zfmisc_1(X34,X35,X36),X37), inference(variable_rename,[status(thm)],[d4_zfmisc_1])).
cnf(c_0_10, negated_conjecture, (esk5_0=X1|~m1_subset_1(X1,esk1_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X4,esk4_0)|esk6_0!=k4_mcart_1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11, plain, (k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (X1=k4_mcart_1(esk7_5(X2,X3,X4,X5,X1),esk8_5(X2,X3,X4,X5,X1),esk9_5(X2,X3,X4,X5,X1),esk10_5(X2,X3,X4,X5,X1))|X5=k1_xboole_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|~m1_subset_1(X1,k4_zfmisc_1(X2,X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, negated_conjecture, (esk5_0=X1|esk6_0!=k4_tarski(k3_mcart_1(X1,X2,X3),X4)|~m1_subset_1(X4,esk4_0)|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk1_0)), inference(rw,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15, plain, (X5=k1_xboole_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k4_tarski(k3_mcart_1(esk7_5(X2,X3,X4,X5,X1),esk8_5(X2,X3,X4,X5,X1),esk9_5(X2,X3,X4,X5,X1)),esk10_5(X2,X3,X4,X5,X1))|~m1_subset_1(X1,k2_zfmisc_1(k3_zfmisc_1(X2,X3,X4),X5))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_13]), c_0_11])).
cnf(c_0_16, plain, (m1_subset_1(esk10_5(X1,X2,X3,X4,X5),X4)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_17, negated_conjecture, (esk7_5(X1,X2,X3,X4,esk6_0)=esk5_0|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|~m1_subset_1(esk10_5(X1,X2,X3,X4,esk6_0),esk4_0)|~m1_subset_1(esk9_5(X1,X2,X3,X4,esk6_0),esk3_0)|~m1_subset_1(esk8_5(X1,X2,X3,X4,esk6_0),esk2_0)|~m1_subset_1(esk7_5(X1,X2,X3,X4,esk6_0),esk1_0)|~m1_subset_1(esk6_0,k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4))), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15])])).
cnf(c_0_18, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|m1_subset_1(esk10_5(X1,X2,X3,X4,X5),X4)|~m1_subset_1(X5,k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4))), inference(rw,[status(thm)],[c_0_16, c_0_13])).
cnf(c_0_19, negated_conjecture, (esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_20, plain, (m1_subset_1(esk9_5(X1,X2,X3,X4,X5),X3)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_21, negated_conjecture, (esk7_5(X1,X2,X3,esk4_0,esk6_0)=esk5_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk9_5(X1,X2,X3,esk4_0,esk6_0),esk3_0)|~m1_subset_1(esk8_5(X1,X2,X3,esk4_0,esk6_0),esk2_0)|~m1_subset_1(esk7_5(X1,X2,X3,esk4_0,esk6_0),esk1_0)|~m1_subset_1(esk6_0,k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])).
cnf(c_0_22, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|m1_subset_1(esk9_5(X1,X2,X3,X4,X5),X3)|~m1_subset_1(X5,k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4))), inference(rw,[status(thm)],[c_0_20, c_0_13])).
cnf(c_0_23, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_24, plain, (m1_subset_1(esk8_5(X1,X2,X3,X4,X5),X2)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_25, plain, ![X38, X39, X40, X41, X42, X43, X44, X45, X46]:((k8_mcart_1(X38,X39,X40,X41,X42)=X43|X42!=k4_mcart_1(X43,X44,X45,X46)|X41=k1_xboole_0|X40=k1_xboole_0|X39=k1_xboole_0|X38=k1_xboole_0|~m1_subset_1(X42,k4_zfmisc_1(X38,X39,X40,X41)))&((k9_mcart_1(X38,X39,X40,X41,X42)=X44|X42!=k4_mcart_1(X43,X44,X45,X46)|X41=k1_xboole_0|X40=k1_xboole_0|X39=k1_xboole_0|X38=k1_xboole_0|~m1_subset_1(X42,k4_zfmisc_1(X38,X39,X40,X41)))&((k10_mcart_1(X38,X39,X40,X41,X42)=X45|X42!=k4_mcart_1(X43,X44,X45,X46)|X41=k1_xboole_0|X40=k1_xboole_0|X39=k1_xboole_0|X38=k1_xboole_0|~m1_subset_1(X42,k4_zfmisc_1(X38,X39,X40,X41)))&(k11_mcart_1(X38,X39,X40,X41,X42)=X46|X42!=k4_mcart_1(X43,X44,X45,X46)|X41=k1_xboole_0|X40=k1_xboole_0|X39=k1_xboole_0|X38=k1_xboole_0|~m1_subset_1(X42,k4_zfmisc_1(X38,X39,X40,X41)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t78_mcart_1])])])])])).
cnf(c_0_26, negated_conjecture, (esk7_5(X1,X2,esk3_0,esk4_0,esk6_0)=esk5_0|X1=k1_xboole_0|X2=k1_xboole_0|~m1_subset_1(esk8_5(X1,X2,esk3_0,esk4_0,esk6_0),esk2_0)|~m1_subset_1(esk7_5(X1,X2,esk3_0,esk4_0,esk6_0),esk1_0)|~m1_subset_1(esk6_0,k2_zfmisc_1(k3_zfmisc_1(X1,X2,esk3_0),esk4_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_19])).
cnf(c_0_27, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|m1_subset_1(esk8_5(X1,X2,X3,X4,X5),X2)|~m1_subset_1(X5,k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4))), inference(rw,[status(thm)],[c_0_24, c_0_13])).
cnf(c_0_28, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_29, plain, (m1_subset_1(esk7_5(X1,X2,X3,X4,X5),X1)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_30, negated_conjecture, (m1_subset_1(esk6_0,k4_zfmisc_1(esk1_0,esk2_0,esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_31, plain, (k8_mcart_1(X1,X2,X3,X4,X5)=X6|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X5!=k4_mcart_1(X6,X7,X8,X9)|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, negated_conjecture, (esk7_5(X1,esk2_0,esk3_0,esk4_0,esk6_0)=esk5_0|X1=k1_xboole_0|~m1_subset_1(esk7_5(X1,esk2_0,esk3_0,esk4_0,esk6_0),esk1_0)|~m1_subset_1(esk6_0,k2_zfmisc_1(k3_zfmisc_1(X1,esk2_0,esk3_0),esk4_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_19]), c_0_23])).
cnf(c_0_33, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|m1_subset_1(esk7_5(X1,X2,X3,X4,X5),X1)|~m1_subset_1(X5,k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4))), inference(rw,[status(thm)],[c_0_29, c_0_13])).
cnf(c_0_34, negated_conjecture, (m1_subset_1(esk6_0,k2_zfmisc_1(k3_zfmisc_1(esk1_0,esk2_0,esk3_0),esk4_0))), inference(rw,[status(thm)],[c_0_30, c_0_13])).
cnf(c_0_35, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_36, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k8_mcart_1(X1,X2,X3,X4,X5)=X6|X5!=k4_tarski(k3_mcart_1(X6,X7,X8),X9)|~m1_subset_1(X5,k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_13]), c_0_11])).
cnf(c_0_37, negated_conjecture, (esk7_5(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)=esk5_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])]), c_0_35]), c_0_19]), c_0_23]), c_0_28])).
cnf(c_0_38, plain, (k8_mcart_1(X1,X2,X3,X4,k4_tarski(k3_mcart_1(X5,X6,X7),X8))=X5|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(k4_tarski(k3_mcart_1(X5,X6,X7),X8),k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4))), inference(er,[status(thm)],[c_0_36])).
cnf(c_0_39, negated_conjecture, (k4_tarski(k3_mcart_1(esk5_0,esk8_5(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0),esk9_5(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)),esk10_5(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))=esk6_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_37]), c_0_34])]), c_0_19]), c_0_23]), c_0_28]), c_0_35])).
cnf(c_0_40, negated_conjecture, (k8_mcart_1(X1,X2,X3,X4,esk6_0)=esk5_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk6_0,k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_41, negated_conjecture, (esk5_0!=k8_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_42, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_34]), c_0_41]), c_0_19]), c_0_23]), c_0_28]), c_0_35]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 43
# Proof object clause steps            : 32
# Proof object formula steps           : 11
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 35
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 16
# Processed clauses                    : 46
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 46
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 26
# ...of the previous two non-trivial   : 24
# Contextual simplify-reflections      : 0
# Paramodulations                      : 21
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 26
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 17
# Current number of unprocessed clauses: 10
# ...number of literals in the above   : 111
# Current number of archived formulas  : 0
# Current number of archived clauses   : 18
# Clause-clause subsumption calls (NU) : 32
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 14
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2664

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.003 s
# Total time               : 0.010 s
# Maximum resident set size: 28684 pages
