# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_zfmisc_1, axiom, ![X1, X2, X3, X4]:k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4), file('data/MPTP/proofs/7959772a7c664bc6b2c6e2c31dd27376.E_input', d4_zfmisc_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/7959772a7c664bc6b2c6e2c31dd27376.E_input', d3_zfmisc_1)).
fof(d4_mcart_1, axiom, ![X1, X2, X3, X4]:k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4), file('data/MPTP/proofs/7959772a7c664bc6b2c6e2c31dd27376.E_input', d4_mcart_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/7959772a7c664bc6b2c6e2c31dd27376.E_input', d3_mcart_1)).
fof(d2_mcart_1, axiom, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>![X2]:(X2=k2_mcart_1(X1)<=>![X3, X4]:(X1=k4_tarski(X3,X4)=>X2=X4))), file('data/MPTP/proofs/7959772a7c664bc6b2c6e2c31dd27376.E_input', d2_mcart_1)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/7959772a7c664bc6b2c6e2c31dd27376.E_input', t7_mcart_1)).
fof(t60_mcart_1, axiom, ![X1, X2, X3, X4]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&~(![X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>X5=k4_mcart_1(k8_mcart_1(X1,X2,X3,X4,X5),k9_mcart_1(X1,X2,X3,X4,X5),k10_mcart_1(X1,X2,X3,X4,X5),k11_mcart_1(X1,X2,X3,X4,X5))))))))), file('data/MPTP/proofs/7959772a7c664bc6b2c6e2c31dd27376.E_input', t60_mcart_1)).
fof(t80_mcart_1, conjecture, ![X1, X2, X3, X4, X5, X6]:(m1_subset_1(X6,k4_zfmisc_1(X1,X2,X3,X4))=>(![X7]:(m1_subset_1(X7,X1)=>![X8]:(m1_subset_1(X8,X2)=>![X9]:(m1_subset_1(X9,X3)=>![X10]:(m1_subset_1(X10,X4)=>(X6=k4_mcart_1(X7,X8,X9,X10)=>X5=X8)))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|(X4=k1_xboole_0|X5=k9_mcart_1(X1,X2,X3,X4,X6))))))), file('data/MPTP/proofs/7959772a7c664bc6b2c6e2c31dd27376.E_input', t80_mcart_1)).
fof(dt_k11_mcart_1, axiom, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>m1_subset_1(k11_mcart_1(X1,X2,X3,X4,X5),X4)), file('data/MPTP/proofs/7959772a7c664bc6b2c6e2c31dd27376.E_input', dt_k11_mcart_1)).
fof(dt_k10_mcart_1, axiom, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>m1_subset_1(k10_mcart_1(X1,X2,X3,X4,X5),X3)), file('data/MPTP/proofs/7959772a7c664bc6b2c6e2c31dd27376.E_input', dt_k10_mcart_1)).
fof(dt_k8_mcart_1, axiom, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>m1_subset_1(k8_mcart_1(X1,X2,X3,X4,X5),X1)), file('data/MPTP/proofs/7959772a7c664bc6b2c6e2c31dd27376.E_input', dt_k8_mcart_1)).
fof(dt_k9_mcart_1, axiom, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))=>m1_subset_1(k9_mcart_1(X1,X2,X3,X4,X5),X2)), file('data/MPTP/proofs/7959772a7c664bc6b2c6e2c31dd27376.E_input', dt_k9_mcart_1)).
fof(t20_mcart_1, axiom, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1)))), file('data/MPTP/proofs/7959772a7c664bc6b2c6e2c31dd27376.E_input', t20_mcart_1)).
fof(c_0_13, plain, ![X11, X12, X13, X14]:k4_zfmisc_1(X11,X12,X13,X14)=k2_zfmisc_1(k3_zfmisc_1(X11,X12,X13),X14), inference(variable_rename,[status(thm)],[d4_zfmisc_1])).
fof(c_0_14, plain, ![X15, X16, X17]:k3_zfmisc_1(X15,X16,X17)=k2_zfmisc_1(k2_zfmisc_1(X15,X16),X17), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_15, plain, ![X18, X19, X20, X21]:k4_mcart_1(X18,X19,X20,X21)=k4_tarski(k3_mcart_1(X18,X19,X20),X21), inference(variable_rename,[status(thm)],[d4_mcart_1])).
fof(c_0_16, plain, ![X22, X23, X24]:k3_mcart_1(X22,X23,X24)=k4_tarski(k4_tarski(X22,X23),X24), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_17, plain, ![X25, X26, X27, X28, X29, X30, X31]:((X28!=k2_mcart_1(X25)|(X25!=k4_tarski(X29,X30)|X28=X30)|X25!=k4_tarski(X26,X27))&((X25=k4_tarski(esk1_2(X25,X31),esk2_2(X25,X31))|X31=k2_mcart_1(X25)|X25!=k4_tarski(X26,X27))&(X31!=esk2_2(X25,X31)|X31=k2_mcart_1(X25)|X25!=k4_tarski(X26,X27)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_mcart_1])])])])])])).
fof(c_0_18, plain, ![X34, X35]:(k1_mcart_1(k4_tarski(X34,X35))=X34&k2_mcart_1(k4_tarski(X34,X35))=X35), inference(variable_rename,[status(thm)],[t7_mcart_1])).
fof(c_0_19, plain, ![X36, X37, X38, X39, X40]:(X36=k1_xboole_0|(X37=k1_xboole_0|(X38=k1_xboole_0|(X39=k1_xboole_0|(~m1_subset_1(X40,k4_zfmisc_1(X36,X37,X38,X39))|X40=k4_mcart_1(k8_mcart_1(X36,X37,X38,X39,X40),k9_mcart_1(X36,X37,X38,X39,X40),k10_mcart_1(X36,X37,X38,X39,X40),k11_mcart_1(X36,X37,X38,X39,X40))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t60_mcart_1])])])])).
cnf(c_0_20, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_24, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6]:(m1_subset_1(X6,k4_zfmisc_1(X1,X2,X3,X4))=>(![X7]:(m1_subset_1(X7,X1)=>![X8]:(m1_subset_1(X8,X2)=>![X9]:(m1_subset_1(X9,X3)=>![X10]:(m1_subset_1(X10,X4)=>(X6=k4_mcart_1(X7,X8,X9,X10)=>X5=X8)))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|(X4=k1_xboole_0|X5=k9_mcart_1(X1,X2,X3,X4,X6)))))))), inference(assume_negation,[status(cth)],[t80_mcart_1])).
cnf(c_0_25, plain, (X1=k4_tarski(esk1_2(X1,X2),esk2_2(X1,X2))|X2=k2_mcart_1(X1)|X1!=k4_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|X5=k4_mcart_1(k8_mcart_1(X1,X2,X3,X4,X5),k9_mcart_1(X1,X2,X3,X4,X5),k10_mcart_1(X1,X2,X3,X4,X5),k11_mcart_1(X1,X2,X3,X4,X5))|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_29, plain, (k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_22, c_0_23])).
fof(c_0_30, negated_conjecture, ![X47, X48, X49, X50]:(m1_subset_1(esk8_0,k4_zfmisc_1(esk3_0,esk4_0,esk5_0,esk6_0))&((~m1_subset_1(X47,esk3_0)|(~m1_subset_1(X48,esk4_0)|(~m1_subset_1(X49,esk5_0)|(~m1_subset_1(X50,esk6_0)|(esk8_0!=k4_mcart_1(X47,X48,X49,X50)|esk7_0=X48)))))&(esk3_0!=k1_xboole_0&(esk4_0!=k1_xboole_0&(esk5_0!=k1_xboole_0&(esk6_0!=k1_xboole_0&esk7_0!=k9_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])])).
cnf(c_0_31, plain, (k4_tarski(esk1_2(k4_tarski(X1,X2),X3),esk2_2(k4_tarski(X1,X2),X3))=k4_tarski(X1,X2)|X3=X2), inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_25]), c_0_26])).
cnf(c_0_32, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X5=k4_tarski(k4_tarski(k4_tarski(k8_mcart_1(X1,X2,X3,X4,X5),k9_mcart_1(X1,X2,X3,X4,X5)),k10_mcart_1(X1,X2,X3,X4,X5)),k11_mcart_1(X1,X2,X3,X4,X5))|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_28]), c_0_29])).
cnf(c_0_33, negated_conjecture, (m1_subset_1(esk8_0,k4_zfmisc_1(esk3_0,esk4_0,esk5_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_34, plain, ![X51, X52, X53, X54, X55]:(~m1_subset_1(X55,k4_zfmisc_1(X51,X52,X53,X54))|m1_subset_1(k11_mcart_1(X51,X52,X53,X54,X55),X54)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k11_mcart_1])])).
cnf(c_0_35, plain, (k4_tarski(esk1_2(X1,X2),esk2_2(X1,X2))=X1|X2=k11_mcart_1(X3,X4,X5,X6,X1)|X3=k1_xboole_0|X4=k1_xboole_0|X5=k1_xboole_0|X6=k1_xboole_0|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X3,X4),X5),X6))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_36, negated_conjecture, (m1_subset_1(esk8_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk3_0,esk4_0),esk5_0),esk6_0))), inference(rw,[status(thm)],[c_0_33, c_0_28])).
cnf(c_0_37, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, negated_conjecture, (esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39, negated_conjecture, (esk5_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40, negated_conjecture, (esk6_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_41, negated_conjecture, (esk7_0=X2|~m1_subset_1(X1,esk3_0)|~m1_subset_1(X2,esk4_0)|~m1_subset_1(X3,esk5_0)|~m1_subset_1(X4,esk6_0)|esk8_0!=k4_mcart_1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_42, plain, (m1_subset_1(k11_mcart_1(X2,X3,X4,X5,X1),X5)|~m1_subset_1(X1,k4_zfmisc_1(X2,X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43, negated_conjecture, (k4_tarski(esk1_2(esk8_0,X1),esk2_2(esk8_0,X1))=esk8_0|X1=k11_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40])).
cnf(c_0_44, negated_conjecture, (esk7_0=X2|esk8_0!=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)|~m1_subset_1(X4,esk6_0)|~m1_subset_1(X3,esk5_0)|~m1_subset_1(X2,esk4_0)|~m1_subset_1(X1,esk3_0)), inference(rw,[status(thm)],[c_0_41, c_0_29])).
cnf(c_0_45, plain, (m1_subset_1(k11_mcart_1(X2,X3,X4,X5,X1),X5)|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4),X5))), inference(rw,[status(thm)],[c_0_42, c_0_28])).
cnf(c_0_46, negated_conjecture, (X1=k11_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0)|esk2_2(esk8_0,X1)=k2_mcart_1(esk8_0)), inference(spm,[status(thm)],[c_0_26, c_0_43])).
fof(c_0_47, plain, ![X61, X62, X63, X64, X65]:(~m1_subset_1(X65,k4_zfmisc_1(X61,X62,X63,X64))|m1_subset_1(k10_mcart_1(X61,X62,X63,X64,X65),X63)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k10_mcart_1])])).
cnf(c_0_48, negated_conjecture, (k9_mcart_1(X1,X2,X3,X4,esk8_0)=esk7_0|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|~m1_subset_1(esk8_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))|~m1_subset_1(k11_mcart_1(X1,X2,X3,X4,esk8_0),esk6_0)|~m1_subset_1(k10_mcart_1(X1,X2,X3,X4,esk8_0),esk5_0)|~m1_subset_1(k9_mcart_1(X1,X2,X3,X4,esk8_0),esk4_0)|~m1_subset_1(k8_mcart_1(X1,X2,X3,X4,esk8_0),esk3_0)), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_32])])).
cnf(c_0_49, negated_conjecture, (esk7_0!=k9_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_50, negated_conjecture, (esk2_2(esk8_0,X1)=k2_mcart_1(esk8_0)|m1_subset_1(X1,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_36])])).
cnf(c_0_51, plain, (m1_subset_1(k10_mcart_1(X2,X3,X4,X5,X1),X4)|~m1_subset_1(X1,k4_zfmisc_1(X2,X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_52, plain, ![X56, X57, X58, X59, X60]:(~m1_subset_1(X60,k4_zfmisc_1(X56,X57,X58,X59))|m1_subset_1(k8_mcart_1(X56,X57,X58,X59,X60),X56)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_mcart_1])])).
cnf(c_0_53, negated_conjecture, (esk2_2(esk8_0,X1)=k2_mcart_1(esk8_0)|~m1_subset_1(k10_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0),esk5_0)|~m1_subset_1(k9_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0),esk4_0)|~m1_subset_1(k8_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0),esk3_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_46]), c_0_36])]), c_0_49]), c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_50])).
cnf(c_0_54, plain, (m1_subset_1(k10_mcart_1(X2,X3,X4,X5,X1),X4)|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4),X5))), inference(rw,[status(thm)],[c_0_51, c_0_28])).
cnf(c_0_55, plain, (m1_subset_1(k8_mcart_1(X2,X3,X4,X5,X1),X2)|~m1_subset_1(X1,k4_zfmisc_1(X2,X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_52])).
fof(c_0_56, plain, ![X66, X67, X68, X69, X70]:(~m1_subset_1(X70,k4_zfmisc_1(X66,X67,X68,X69))|m1_subset_1(k9_mcart_1(X66,X67,X68,X69,X70),X67)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k9_mcart_1])])).
cnf(c_0_57, negated_conjecture, (esk2_2(esk8_0,X1)=k2_mcart_1(esk8_0)|~m1_subset_1(k9_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0),esk4_0)|~m1_subset_1(k8_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_36])])).
cnf(c_0_58, plain, (m1_subset_1(k8_mcart_1(X2,X3,X4,X5,X1),X2)|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4),X5))), inference(rw,[status(thm)],[c_0_55, c_0_28])).
cnf(c_0_59, plain, (m1_subset_1(k9_mcart_1(X2,X3,X4,X5,X1),X3)|~m1_subset_1(X1,k4_zfmisc_1(X2,X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_60, negated_conjecture, (esk2_2(esk8_0,X1)=k2_mcart_1(esk8_0)|~m1_subset_1(k9_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0),esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_36])])).
cnf(c_0_61, plain, (m1_subset_1(k9_mcart_1(X2,X3,X4,X5,X1),X3)|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4),X5))), inference(rw,[status(thm)],[c_0_59, c_0_28])).
cnf(c_0_62, negated_conjecture, (esk2_2(esk8_0,X1)=k2_mcart_1(esk8_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_36])])).
fof(c_0_63, plain, ![X71, X72, X73]:((X71!=k1_mcart_1(X71)|X71!=k4_tarski(X72,X73))&(X71!=k2_mcart_1(X71)|X71!=k4_tarski(X72,X73))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t20_mcart_1])])])])])).
cnf(c_0_64, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_65, negated_conjecture, (X1=k11_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0)|k4_tarski(esk1_2(esk8_0,X1),k2_mcart_1(esk8_0))=esk8_0), inference(rw,[status(thm)],[c_0_43, c_0_62])).
cnf(c_0_66, plain, (X1!=k2_mcart_1(X1)|X1!=k4_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_67, negated_conjecture, (X1=k11_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0)|esk1_2(esk8_0,X1)=k1_mcart_1(esk8_0)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_68, plain, (k4_tarski(X1,X2)!=X2), inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_66]), c_0_26])).
cnf(c_0_69, negated_conjecture, (esk1_2(esk8_0,X1)=k1_mcart_1(esk8_0)|m1_subset_1(X1,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_67]), c_0_36])])).
cnf(c_0_70, plain, (X1=X2|esk2_2(k4_tarski(X3,X2),X1)!=k4_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_68, c_0_31])).
cnf(c_0_71, plain, (k4_tarski(esk1_2(k4_tarski(X1,X2),X3),esk2_2(k4_tarski(X1,X2),X3))=k4_tarski(X1,X2)|X3=esk2_2(k4_tarski(X1,X2),X4)|X4=X2), inference(spm,[status(thm)],[c_0_31, c_0_31])).
cnf(c_0_72, negated_conjecture, (esk1_2(esk8_0,X1)=k1_mcart_1(esk8_0)|~m1_subset_1(k10_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0),esk5_0)|~m1_subset_1(k9_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0),esk4_0)|~m1_subset_1(k8_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0),esk3_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_67]), c_0_36])]), c_0_49]), c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_69])).
cnf(c_0_73, plain, (k4_tarski(esk1_2(k4_tarski(X1,X2),k4_tarski(X1,X2)),esk2_2(k4_tarski(X1,X2),k4_tarski(X1,X2)))=k4_tarski(X1,X2)|X3=X2), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71])])).
cnf(c_0_74, plain, (esk1_2(k4_tarski(X1,X2),X3)=X1|X3=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_31]), c_0_64])).
cnf(c_0_75, negated_conjecture, (esk1_2(esk8_0,X1)=k1_mcart_1(esk8_0)|~m1_subset_1(k9_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0),esk4_0)|~m1_subset_1(k8_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_54]), c_0_36])])).
cnf(c_0_76, plain, (esk2_2(k4_tarski(X1,X2),k4_tarski(X1,X2))=X2|X3=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_73]), c_0_26])).
cnf(c_0_77, plain, (esk1_2(X1,X2)=k4_tarski(k4_tarski(k8_mcart_1(X3,X4,X5,X6,X1),k9_mcart_1(X3,X4,X5,X6,X1)),k10_mcart_1(X3,X4,X5,X6,X1))|X2=k11_mcart_1(X3,X4,X5,X6,X1)|X3=k1_xboole_0|X4=k1_xboole_0|X5=k1_xboole_0|X6=k1_xboole_0|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X3,X4),X5),X6))), inference(spm,[status(thm)],[c_0_74, c_0_32])).
cnf(c_0_78, negated_conjecture, (esk1_2(esk8_0,X1)=k1_mcart_1(esk8_0)|~m1_subset_1(k9_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0),esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_58]), c_0_36])])).
cnf(c_0_79, plain, (esk2_2(k4_tarski(X1,X2),k4_tarski(X1,X2))=X2), inference(ef,[status(thm)],[c_0_76])).
cnf(c_0_80, plain, (k4_tarski(k4_tarski(k8_mcart_1(X1,X2,X3,X4,X5),k9_mcart_1(X1,X2,X3,X4,X5)),k10_mcart_1(X1,X2,X3,X4,X5))=k1_mcart_1(X5)|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(spm,[status(thm)],[c_0_64, c_0_32])).
cnf(c_0_81, plain, (k2_mcart_1(esk1_2(X1,X2))=k10_mcart_1(X3,X4,X5,X6,X1)|X2=k11_mcart_1(X3,X4,X5,X6,X1)|X6=k1_xboole_0|X5=k1_xboole_0|X4=k1_xboole_0|X3=k1_xboole_0|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X3,X4),X5),X6))), inference(spm,[status(thm)],[c_0_26, c_0_77])).
cnf(c_0_82, negated_conjecture, (esk1_2(esk8_0,X1)=k1_mcart_1(esk8_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_61]), c_0_36])])).
cnf(c_0_83, plain, (X1!=k1_mcart_1(X1)|X1!=k4_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_84, plain, (X1=esk2_2(k4_tarski(X2,X3),X4)|esk2_2(k4_tarski(X2,X3),X1)=X3|X4=X3), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_71]), c_0_79])).
cnf(c_0_85, plain, (k4_tarski(esk1_2(k1_mcart_1(X1),X2),esk2_2(k1_mcart_1(X1),X2))=k1_mcart_1(X1)|X2=k10_mcart_1(X3,X4,X5,X6,X1)|X6=k1_xboole_0|X5=k1_xboole_0|X4=k1_xboole_0|X3=k1_xboole_0|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X3,X4),X5),X6))), inference(spm,[status(thm)],[c_0_31, c_0_80])).
cnf(c_0_86, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|k11_mcart_1(X1,X2,X3,X4,X5)!=X5|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(spm,[status(thm)],[c_0_68, c_0_32])).
cnf(c_0_87, negated_conjecture, (k2_mcart_1(k1_mcart_1(esk8_0))=k10_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0)|X1=k11_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_36]), c_0_82]), c_0_40]), c_0_39]), c_0_38]), c_0_37])).
cnf(c_0_88, plain, (k4_tarski(X1,X2)!=X1), inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_83]), c_0_64])).
cnf(c_0_89, plain, (esk2_2(k4_tarski(X1,X2),k4_tarski(X3,esk2_2(k4_tarski(X1,X2),X4)))=X2|X4=X2), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_84])])).
cnf(c_0_90, negated_conjecture, (k4_tarski(esk1_2(k1_mcart_1(esk8_0),X1),esk2_2(k1_mcart_1(esk8_0),X1))=k1_mcart_1(esk8_0)|X1=k10_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_36]), c_0_40]), c_0_39]), c_0_38]), c_0_37])).
cnf(c_0_91, negated_conjecture, (k2_mcart_1(k1_mcart_1(esk8_0))=k10_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0)), inference(er,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_36])]), c_0_37]), c_0_38]), c_0_39]), c_0_40])])).
cnf(c_0_92, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|k10_mcart_1(X4,X3,X2,X1,X5)!=k1_mcart_1(X5)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X3),X2),X1))), inference(spm,[status(thm)],[c_0_68, c_0_80])).
cnf(c_0_93, plain, (k2_mcart_1(k1_mcart_1(X1))=k10_mcart_1(X2,X3,X4,X5,X1)|X5=k1_xboole_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4),X5))), inference(spm,[status(thm)],[c_0_26, c_0_80])).
cnf(c_0_94, plain, (esk2_2(k4_tarski(X1,X2),k4_tarski(esk2_2(k4_tarski(X1,X2),X3),X4))=X2|X3=X2), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_84])])).
cnf(c_0_95, plain, (esk2_2(k4_tarski(X1,X2),k4_tarski(X3,X2))=X2), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_79]), c_0_68])).
cnf(c_0_96, negated_conjecture, (X1=k11_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0)|k4_tarski(k1_mcart_1(esk8_0),k2_mcart_1(esk8_0))=esk8_0), inference(rw,[status(thm)],[c_0_65, c_0_82])).
cnf(c_0_97, negated_conjecture, (esk2_2(k1_mcart_1(esk8_0),X1)=k10_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0)|X1=k10_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_90]), c_0_91])).
cnf(c_0_98, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|k2_mcart_1(k1_mcart_1(X5))!=k1_mcart_1(X5)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_99, plain, (esk2_2(k4_tarski(X1,X2),k4_tarski(X2,X3))=X2), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_68])).
cnf(c_0_100, negated_conjecture, (k4_tarski(k1_mcart_1(esk8_0),k2_mcart_1(esk8_0))=esk8_0), inference(er,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_96]), c_0_36])]), c_0_37]), c_0_38]), c_0_39]), c_0_40])])).
cnf(c_0_101, negated_conjecture, (k4_tarski(esk1_2(k1_mcart_1(esk8_0),X1),k10_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0))=k1_mcart_1(esk8_0)|X1=k10_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0)), inference(spm,[status(thm)],[c_0_90, c_0_97])).
cnf(c_0_102, negated_conjecture, (X1=k10_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0)|esk1_2(k1_mcart_1(esk8_0),X1)=k1_mcart_1(k1_mcart_1(esk8_0))), inference(spm,[status(thm)],[c_0_64, c_0_90])).
cnf(c_0_103, negated_conjecture, (k2_mcart_1(k1_mcart_1(esk8_0))!=k1_mcart_1(esk8_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40])).
cnf(c_0_104, plain, (k4_tarski(esk1_2(k4_tarski(X1,X2),k4_tarski(X3,X2)),X2)=k4_tarski(X1,X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_95]), c_0_68])).
cnf(c_0_105, negated_conjecture, (esk2_2(k4_tarski(X1,k1_mcart_1(esk8_0)),esk8_0)=k1_mcart_1(esk8_0)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_106, negated_conjecture, (k4_tarski(k1_mcart_1(k1_mcart_1(esk8_0)),k10_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0))=k1_mcart_1(esk8_0)|X1=k10_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0)), inference(spm,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_107, negated_conjecture, (k10_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0)!=k1_mcart_1(esk8_0)), inference(rw,[status(thm)],[c_0_103, c_0_91])).
cnf(c_0_108, plain, (esk1_2(k4_tarski(X1,X2),k4_tarski(X3,X2))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_104]), c_0_64])).
cnf(c_0_109, negated_conjecture, (k4_tarski(k1_mcart_1(k1_mcart_1(esk8_0)),k10_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0))=k1_mcart_1(esk8_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_107])).
cnf(c_0_110, negated_conjecture, (esk1_2(k4_tarski(X1,k10_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0)),k1_mcart_1(esk8_0))=X1), inference(spm,[status(thm)],[c_0_108, c_0_109])).
cnf(c_0_111, negated_conjecture, (esk1_2(k1_mcart_1(esk8_0),k1_mcart_1(esk8_0))=k1_mcart_1(k1_mcart_1(esk8_0))), inference(spm,[status(thm)],[c_0_110, c_0_109])).
cnf(c_0_112, negated_conjecture, (k4_tarski(k8_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0),k9_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0))=k1_mcart_1(k1_mcart_1(esk8_0))), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_80]), c_0_36])]), c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_111])).
cnf(c_0_113, plain, (k11_mcart_1(X1,X2,X3,X4,X5)=k2_mcart_1(X5)|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(spm,[status(thm)],[c_0_26, c_0_32])).
cnf(c_0_114, negated_conjecture, (k4_tarski(k1_mcart_1(esk8_0),k11_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0))=esk8_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_112]), c_0_109]), c_0_36])]), c_0_40]), c_0_39]), c_0_38]), c_0_37])).
cnf(c_0_115, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|m1_subset_1(k2_mcart_1(X5),X1)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X3),X2),X1))), inference(spm,[status(thm)],[c_0_45, c_0_113])).
cnf(c_0_116, negated_conjecture, (k11_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0)=k2_mcart_1(esk8_0)), inference(spm,[status(thm)],[c_0_26, c_0_114])).
cnf(c_0_117, negated_conjecture, (m1_subset_1(k2_mcart_1(esk8_0),esk6_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_36]), c_0_40]), c_0_39]), c_0_38]), c_0_37])).
cnf(c_0_118, negated_conjecture, (~m1_subset_1(k10_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0),esk5_0)|~m1_subset_1(k9_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0),esk4_0)|~m1_subset_1(k8_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0),esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_116]), c_0_36]), c_0_117])]), c_0_49]), c_0_37]), c_0_38]), c_0_39]), c_0_40])).
cnf(c_0_119, negated_conjecture, (~m1_subset_1(k9_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0),esk4_0)|~m1_subset_1(k8_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_54]), c_0_36])])).
cnf(c_0_120, negated_conjecture, (~m1_subset_1(k9_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0,esk8_0),esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_58]), c_0_36])])).
cnf(c_0_121, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_61]), c_0_36])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 122
# Proof object clause steps            : 95
# Proof object formula steps           : 27
# Proof object conjectures             : 50
# Proof object clause conjectures      : 47
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 59
# Proof object simplifying inferences  : 127
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 19
# Processed clauses                    : 4899
# ...of these trivial                  : 74
# ...subsumed                          : 3986
# ...remaining for further processing  : 839
# Other redundant clauses eliminated   : 21720
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 80
# Backward-rewritten                   : 90
# Generated clauses                    : 284741
# ...of the previous two non-trivial   : 258953
# Contextual simplify-reflections      : 3
# Paramodulations                      : 262719
# Factorizations                       : 294
# Equation resolutions                 : 21729
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 645
#    Positive orientable unit clauses  : 34
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 14
#    Non-unit-clauses                  : 597
# Current number of unprocessed clauses: 251307
# ...number of literals in the above   : 1903662
# Current number of archived formulas  : 0
# Current number of archived clauses   : 193
# Clause-clause subsumption calls (NU) : 148561
# Rec. Clause-clause subsumption calls : 30656
# Non-unit clause-clause subsumptions  : 1835
# Unit Clause-clause subsumption calls : 5572
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 121
# BW rewrite match successes           : 49
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6238986

# -------------------------------------------------
# User time                : 5.279 s
# System time              : 0.100 s
# Total time               : 5.379 s
# Maximum resident set size: 210352 pages
