# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', t1_subset)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', d3_tarski)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', t4_xboole_0)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', t1_xboole_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', t2_xboole_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', t3_subset)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', d10_xboole_0)).
fof(t3_xboole_1, axiom, ![X1]:(r1_tarski(X1,k1_xboole_0)=>X1=k1_xboole_0), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', t3_xboole_1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', t7_ordinal1)).
fof(fraenkel_a_2_5_lattice3, axiom, ![X1, X2, X3]:((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))=>(r2_hidden(X1,a_2_5_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r3_lattices(X2,X4,X5)&r2_hidden(X5,X3))))))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', fraenkel_a_2_5_lattice3)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', t65_xboole_1)).
fof(t50_lattice3, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))&k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', t50_lattice3)).
fof(d21_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X3=k15_lattice3(X1,X2)<=>(r4_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r4_lattice3(X1,X4,X2)=>r1_lattices(X1,X3,X4)))))))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', d21_lattice3)).
fof(dt_k15_lattice3, axiom, ![X1, X2]:((~(v2_struct_0(X1))&l3_lattices(X1))=>m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', dt_k15_lattice3)).
fof(t47_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(k15_lattice3(X1,X2)=k15_lattice3(X1,a_2_5_lattice3(X1,X2))&k16_lattice3(X1,X2)=k16_lattice3(X1,a_2_6_lattice3(X1,X2)))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', t47_lattice3)).
fof(t26_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&(v4_lattices(X1)&l2_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattices(X1,X2,X3)&r1_lattices(X1,X3,X2))=>X2=X3)))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', t26_lattices)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', cc1_lattices)).
fof(dt_l3_lattices, axiom, ![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', dt_l3_lattices)).
fof(t45_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(X2=k15_lattice3(X1,a_2_3_lattice3(X1,X2))&X2=k16_lattice3(X1,a_2_4_lattice3(X1,X2))))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', t45_lattice3)).
fof(d17_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r4_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X4,X2)))))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', d17_lattice3)).
fof(fraenkel_a_2_3_lattice3, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))&m1_subset_1(X3,u1_struct_0(X2)))=>(r2_hidden(X1,a_2_3_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&r3_lattices(X2,X4,X3))))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', fraenkel_a_2_3_lattice3)).
fof(t46_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(r1_tarski(X2,X3)=>(r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3))&r3_lattices(X1,k16_lattice3(X1,X3),k16_lattice3(X1,X2))))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', t46_lattice3)).
fof(d9_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>(v9_lattices(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_lattices(X1,X2,k1_lattices(X1,X2,X3))=X2)))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', d9_lattices)).
fof(d3_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l2_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattices(X1,X2,X3)<=>k1_lattices(X1,X2,X3)=X3)))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', d3_lattices)).
fof(d16_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v13_lattices(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(X2=k5_lattices(X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_lattices(X1,X2,X3)=X2&k2_lattices(X1,X3,X2)=X2)))))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', d16_lattices)).
fof(dt_k5_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>m1_subset_1(k5_lattices(X1),u1_struct_0(X1))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', dt_k5_lattices)).
fof(d6_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v6_lattices(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2))))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', d6_lattices)).
fof(d13_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v13_lattices(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_lattices(X1,X2,X3)=X2&k2_lattices(X1,X3,X2)=X2))))), file('data/MPTP/proofs/7a042822fd5a4109af1a561907648a00.E_input', d13_lattices)).
fof(c_0_28, plain, ![X71, X72]:(~r2_hidden(X71,X72)|m1_subset_1(X71,X72)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
fof(c_0_29, plain, ![X40, X41, X42, X43, X44]:((~r1_tarski(X40,X41)|(~r2_hidden(X42,X40)|r2_hidden(X42,X41)))&((r2_hidden(esk7_2(X43,X44),X43)|r1_tarski(X43,X44))&(~r2_hidden(esk7_2(X43,X44),X44)|r1_tarski(X43,X44)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_30, plain, ![X8, X9, X11, X12, X13]:((r1_xboole_0(X8,X9)|r2_hidden(esk2_2(X8,X9),k3_xboole_0(X8,X9)))&(~r2_hidden(X13,k3_xboole_0(X11,X12))|~r1_xboole_0(X11,X12))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_31, plain, ![X22, X23, X24]:(~r1_tarski(X22,X23)|~r1_tarski(X23,X24)|r1_tarski(X22,X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_32, plain, ![X25]:r1_tarski(k1_xboole_0,X25), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_33, plain, ![X38, X39]:((~m1_subset_1(X38,k1_zfmisc_1(X39))|r1_tarski(X38,X39))&(~r1_tarski(X38,X39)|m1_subset_1(X38,k1_zfmisc_1(X39)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_34, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, plain, (r2_hidden(esk7_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_36, plain, ![X69, X70]:(((r1_tarski(X69,X70)|X69!=X70)&(r1_tarski(X70,X69)|X69!=X70))&(~r1_tarski(X69,X70)|~r1_tarski(X70,X69)|X69=X70)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_37, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_39, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_41, plain, (r1_tarski(X1,X2)|m1_subset_1(esk7_2(X1,X2),X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_42, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_43, plain, (r1_tarski(k3_xboole_0(X1,X2),X3)|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_37, c_0_35])).
cnf(c_0_44, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_45, plain, (r1_tarski(esk7_2(k1_zfmisc_1(X1),X2),X1)|r1_tarski(k1_zfmisc_1(X1),X2)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
fof(c_0_46, plain, ![X74]:(~r1_tarski(X74,k1_xboole_0)|X74=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_xboole_1])])).
fof(c_0_47, plain, ![X36, X37]:(~r2_hidden(X36,X37)|~r1_tarski(X37,X36)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_48, plain, (X1=k3_xboole_0(X2,X3)|~r1_tarski(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_49, plain, (r1_tarski(esk7_2(k1_zfmisc_1(k1_xboole_0),X1),X2)|r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_50, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_51, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_52, plain, ![X49, X50, X51, X54, X55, X56]:(((m1_subset_1(esk8_3(X49,X50,X51),u1_struct_0(X50))|~r2_hidden(X49,a_2_5_lattice3(X50,X51))|(v2_struct_0(X50)|(~v10_lattices(X50)|(~v4_lattice3(X50)|~l3_lattices(X50)))))&((X49=esk8_3(X49,X50,X51)|~r2_hidden(X49,a_2_5_lattice3(X50,X51))|(v2_struct_0(X50)|(~v10_lattices(X50)|(~v4_lattice3(X50)|~l3_lattices(X50)))))&((m1_subset_1(esk9_3(X49,X50,X51),u1_struct_0(X50))|~r2_hidden(X49,a_2_5_lattice3(X50,X51))|(v2_struct_0(X50)|(~v10_lattices(X50)|(~v4_lattice3(X50)|~l3_lattices(X50)))))&((r3_lattices(X50,esk8_3(X49,X50,X51),esk9_3(X49,X50,X51))|~r2_hidden(X49,a_2_5_lattice3(X50,X51))|(v2_struct_0(X50)|(~v10_lattices(X50)|(~v4_lattice3(X50)|~l3_lattices(X50)))))&(r2_hidden(esk9_3(X49,X50,X51),X51)|~r2_hidden(X49,a_2_5_lattice3(X50,X51))|(v2_struct_0(X50)|(~v10_lattices(X50)|(~v4_lattice3(X50)|~l3_lattices(X50)))))))))&(~m1_subset_1(X55,u1_struct_0(X50))|(X49!=X55|(~m1_subset_1(X56,u1_struct_0(X50))|(~r3_lattices(X50,X55,X56)|~r2_hidden(X56,X54))))|r2_hidden(X49,a_2_5_lattice3(X50,X54))|(v2_struct_0(X50)|(~v10_lattices(X50)|(~v4_lattice3(X50)|~l3_lattices(X50)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_5_lattice3])])])])])])])).
cnf(c_0_53, plain, (esk7_2(k1_zfmisc_1(k1_xboole_0),X1)=k3_xboole_0(X2,X3)|r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)|~r1_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
fof(c_0_54, plain, ![X78]:r1_xboole_0(X78,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_55, plain, (esk7_2(k1_zfmisc_1(k1_xboole_0),X1)=k1_xboole_0|r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_50, c_0_45])).
cnf(c_0_56, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_51, c_0_39])).
cnf(c_0_57, plain, (r2_hidden(esk9_3(X1,X2,X3),X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_5_lattice3(X2,X3))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_58, plain, (r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)|r2_hidden(k3_xboole_0(X2,X3),k1_zfmisc_1(k1_xboole_0))|~r1_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_35, c_0_53])).
cnf(c_0_59, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_60, plain, (r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_35, c_0_55])).
cnf(c_0_61, plain, (v2_struct_0(X1)|~r2_hidden(X2,a_2_5_lattice3(X1,k1_xboole_0))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_62, plain, (r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)|r2_hidden(k3_xboole_0(X2,k1_xboole_0),k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_63, plain, (k3_xboole_0(X1,X2)=k1_zfmisc_1(k1_xboole_0)|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_48, c_0_60])).
cnf(c_0_64, plain, (r1_tarski(a_2_5_lattice3(X1,k1_xboole_0),X2)|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_61, c_0_35])).
cnf(c_0_65, plain, (r2_hidden(k3_xboole_0(X1,k1_xboole_0),k1_zfmisc_1(k1_xboole_0))|~r2_hidden(X2,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_51, c_0_62])).
cnf(c_0_66, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_zfmisc_1(k1_xboole_0)|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_63, c_0_59])).
cnf(c_0_67, plain, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_50, c_0_60])).
cnf(c_0_68, plain, (X1=a_2_5_lattice3(X2,k1_xboole_0)|v2_struct_0(X2)|~r1_tarski(X1,a_2_5_lattice3(X2,k1_xboole_0))|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)), inference(spm,[status(thm)],[c_0_42, c_0_64])).
fof(c_0_69, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))&k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0)))), inference(assume_negation,[status(cth)],[t50_lattice3])).
cnf(c_0_70, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_zfmisc_1(k1_xboole_0)|r2_hidden(k3_xboole_0(X2,k1_xboole_0),k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_71, plain, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|r2_hidden(k3_xboole_0(X1,k1_xboole_0),k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_65, c_0_67])).
cnf(c_0_72, plain, (esk7_2(k1_zfmisc_1(k1_xboole_0),X1)=a_2_5_lattice3(X2,k1_xboole_0)|r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)|v2_struct_0(X2)|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)), inference(spm,[status(thm)],[c_0_68, c_0_49])).
fof(c_0_73, negated_conjecture, ((~v2_struct_0(esk1_0)&(v10_lattices(esk1_0)&(v4_lattice3(esk1_0)&l3_lattices(esk1_0))))&(v2_struct_0(esk1_0)|(~v10_lattices(esk1_0)|(~v13_lattices(esk1_0)|~l3_lattices(esk1_0)))|k5_lattices(esk1_0)!=k15_lattice3(esk1_0,k1_xboole_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_69])])])])).
cnf(c_0_74, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_zfmisc_1(k1_xboole_0)|m1_subset_1(k3_xboole_0(X2,k1_xboole_0),k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_34, c_0_70])).
cnf(c_0_75, plain, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|m1_subset_1(k3_xboole_0(X1,k1_xboole_0),k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_34, c_0_71])).
cnf(c_0_76, plain, (r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)|r2_hidden(a_2_5_lattice3(X2,k1_xboole_0),k1_zfmisc_1(k1_xboole_0))|v2_struct_0(X2)|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)), inference(spm,[status(thm)],[c_0_35, c_0_72])).
cnf(c_0_77, negated_conjecture, (l3_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_78, negated_conjecture, (v4_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_79, negated_conjecture, (v10_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_80, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_81, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_zfmisc_1(k1_xboole_0)|r1_tarski(k3_xboole_0(X2,k1_xboole_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_40, c_0_74])).
cnf(c_0_82, plain, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|r1_tarski(k3_xboole_0(X1,k1_xboole_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_40, c_0_75])).
cnf(c_0_83, negated_conjecture, (r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)|r2_hidden(a_2_5_lattice3(esk1_0,k1_xboole_0),k1_zfmisc_1(k1_xboole_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78]), c_0_79])]), c_0_80])).
cnf(c_0_84, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_zfmisc_1(k1_xboole_0)|k3_xboole_0(X2,k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_81]), c_0_39])])).
cnf(c_0_85, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0|k1_zfmisc_1(k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_82]), c_0_39])])).
cnf(c_0_86, negated_conjecture, (r2_hidden(a_2_5_lattice3(esk1_0,k1_xboole_0),k1_zfmisc_1(k1_xboole_0))|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_51, c_0_83])).
cnf(c_0_87, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(csr,[status(thm)],[inference(ef,[status(thm)],[c_0_84]), c_0_85])).
fof(c_0_88, plain, ![X89, X90, X91, X92]:(((r4_lattice3(X89,X91,X90)|X91!=k15_lattice3(X89,X90)|~m1_subset_1(X91,u1_struct_0(X89))|(v2_struct_0(X89)|(~v10_lattices(X89)|(~v4_lattice3(X89)|~l3_lattices(X89))))|(v2_struct_0(X89)|~l3_lattices(X89)))&(~m1_subset_1(X92,u1_struct_0(X89))|(~r4_lattice3(X89,X92,X90)|r1_lattices(X89,X91,X92))|X91!=k15_lattice3(X89,X90)|~m1_subset_1(X91,u1_struct_0(X89))|(v2_struct_0(X89)|(~v10_lattices(X89)|(~v4_lattice3(X89)|~l3_lattices(X89))))|(v2_struct_0(X89)|~l3_lattices(X89))))&((m1_subset_1(esk15_3(X89,X90,X91),u1_struct_0(X89))|~r4_lattice3(X89,X91,X90)|X91=k15_lattice3(X89,X90)|~m1_subset_1(X91,u1_struct_0(X89))|(v2_struct_0(X89)|(~v10_lattices(X89)|(~v4_lattice3(X89)|~l3_lattices(X89))))|(v2_struct_0(X89)|~l3_lattices(X89)))&((r4_lattice3(X89,esk15_3(X89,X90,X91),X90)|~r4_lattice3(X89,X91,X90)|X91=k15_lattice3(X89,X90)|~m1_subset_1(X91,u1_struct_0(X89))|(v2_struct_0(X89)|(~v10_lattices(X89)|(~v4_lattice3(X89)|~l3_lattices(X89))))|(v2_struct_0(X89)|~l3_lattices(X89)))&(~r1_lattices(X89,X91,esk15_3(X89,X90,X91))|~r4_lattice3(X89,X91,X90)|X91=k15_lattice3(X89,X90)|~m1_subset_1(X91,u1_struct_0(X89))|(v2_struct_0(X89)|(~v10_lattices(X89)|(~v4_lattice3(X89)|~l3_lattices(X89))))|(v2_struct_0(X89)|~l3_lattices(X89)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d21_lattice3])])])])])])).
cnf(c_0_89, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|r2_hidden(a_2_5_lattice3(esk1_0,k1_xboole_0),k1_zfmisc_1(k1_xboole_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_66]), c_0_87])).
cnf(c_0_90, plain, (r1_lattices(X2,X4,X1)|v2_struct_0(X2)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r4_lattice3(X2,X1,X3)|X4!=k15_lattice3(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_88])).
fof(c_0_91, plain, ![X28, X29]:(v2_struct_0(X28)|~l3_lattices(X28)|m1_subset_1(k15_lattice3(X28,X29),u1_struct_0(X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k15_lattice3])])])).
cnf(c_0_92, plain, (X1=k1_zfmisc_1(k1_xboole_0)|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))|~r1_tarski(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_42, c_0_60])).
cnf(c_0_93, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|m1_subset_1(a_2_5_lattice3(esk1_0,k1_xboole_0),k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_34, c_0_89])).
cnf(c_0_94, plain, (v2_struct_0(X2)|r1_lattices(X2,X4,X1)|X4!=k15_lattice3(X2,X3)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)|~r4_lattice3(X2,X1,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))), inference(cn,[status(thm)],[c_0_90])).
cnf(c_0_95, plain, (v2_struct_0(X1)|m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
fof(c_0_96, plain, ![X57, X58]:((k15_lattice3(X57,X58)=k15_lattice3(X57,a_2_5_lattice3(X57,X58))|(v2_struct_0(X57)|(~v10_lattices(X57)|(~v4_lattice3(X57)|~l3_lattices(X57)))))&(k16_lattice3(X57,X58)=k16_lattice3(X57,a_2_6_lattice3(X57,X58))|(v2_struct_0(X57)|(~v10_lattices(X57)|(~v4_lattice3(X57)|~l3_lattices(X57)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t47_lattice3])])])])])).
cnf(c_0_97, plain, (a_2_5_lattice3(X1,k1_xboole_0)=k1_zfmisc_1(k1_xboole_0)|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_92, c_0_64])).
cnf(c_0_98, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|r1_tarski(a_2_5_lattice3(esk1_0,k1_xboole_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_40, c_0_93])).
fof(c_0_99, plain, ![X14, X15, X16]:(v2_struct_0(X14)|(~v4_lattices(X14)|~l2_lattices(X14))|(~m1_subset_1(X15,u1_struct_0(X14))|(~m1_subset_1(X16,u1_struct_0(X14))|(~r1_lattices(X14,X15,X16)|~r1_lattices(X14,X16,X15)|X15=X16)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t26_lattices])])])])).
fof(c_0_100, plain, ![X30]:((~v2_struct_0(X30)|(v2_struct_0(X30)|~v10_lattices(X30))|~l3_lattices(X30))&((v4_lattices(X30)|(v2_struct_0(X30)|~v10_lattices(X30))|~l3_lattices(X30))&((v5_lattices(X30)|(v2_struct_0(X30)|~v10_lattices(X30))|~l3_lattices(X30))&((v6_lattices(X30)|(v2_struct_0(X30)|~v10_lattices(X30))|~l3_lattices(X30))&((v7_lattices(X30)|(v2_struct_0(X30)|~v10_lattices(X30))|~l3_lattices(X30))&((v8_lattices(X30)|(v2_struct_0(X30)|~v10_lattices(X30))|~l3_lattices(X30))&(v9_lattices(X30)|(v2_struct_0(X30)|~v10_lattices(X30))|~l3_lattices(X30)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
fof(c_0_101, plain, ![X73]:((l1_lattices(X73)|~l3_lattices(X73))&(l2_lattices(X73)|~l3_lattices(X73))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
cnf(c_0_102, plain, (r1_lattices(X1,k15_lattice3(X1,X2),X3)|v2_struct_0(X1)|~r4_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_94]), c_0_95])).
cnf(c_0_103, plain, (k15_lattice3(X1,X2)=k15_lattice3(X1,a_2_5_lattice3(X1,X2))|v2_struct_0(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_104, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)=a_2_5_lattice3(esk1_0,k1_xboole_0)|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_77]), c_0_78]), c_0_79])]), c_0_80])).
cnf(c_0_105, negated_conjecture, (a_2_5_lattice3(esk1_0,k1_xboole_0)=k1_xboole_0|k1_zfmisc_1(k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_98]), c_0_39])])).
cnf(c_0_106, plain, (r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X1)|X2!=k15_lattice3(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_107, plain, (v2_struct_0(X1)|X2=X3|~v4_lattices(X1)|~l2_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattices(X1,X2,X3)|~r1_lattices(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_108, plain, (v4_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_109, plain, (l2_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_110, plain, (r1_lattices(X1,k15_lattice3(X1,X2),X3)|v2_struct_0(X1)|~r4_lattice3(X1,X3,a_2_5_lattice3(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_111, negated_conjecture, (a_2_5_lattice3(esk1_0,k1_xboole_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_56])).
cnf(c_0_112, plain, (v2_struct_0(X1)|r4_lattice3(X1,X2,X3)|X2!=k15_lattice3(X1,X3)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_106])).
fof(c_0_113, plain, ![X26, X27]:((X27=k15_lattice3(X26,a_2_3_lattice3(X26,X27))|~m1_subset_1(X27,u1_struct_0(X26))|(v2_struct_0(X26)|(~v10_lattices(X26)|(~v4_lattice3(X26)|~l3_lattices(X26)))))&(X27=k16_lattice3(X26,a_2_4_lattice3(X26,X27))|~m1_subset_1(X27,u1_struct_0(X26))|(v2_struct_0(X26)|(~v10_lattices(X26)|(~v4_lattice3(X26)|~l3_lattices(X26)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t45_lattice3])])])])])).
cnf(c_0_114, plain, (X1=k15_lattice3(X2,X3)|v2_struct_0(X2)|~r4_lattice3(X2,X1,X3)|~r1_lattices(X2,X1,k15_lattice3(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_102]), c_0_95]), c_0_108]), c_0_109])).
cnf(c_0_115, negated_conjecture, (r1_lattices(esk1_0,k15_lattice3(esk1_0,k1_xboole_0),X1)|~r4_lattice3(esk1_0,X1,k1_xboole_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_77]), c_0_78]), c_0_79])]), c_0_80])).
cnf(c_0_116, plain, (r4_lattice3(X1,k15_lattice3(X1,X2),X2)|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_112]), c_0_95])).
cnf(c_0_117, plain, (X1=k15_lattice3(X2,a_2_3_lattice3(X2,X1))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_113])).
fof(c_0_118, plain, ![X79, X80, X81, X82, X83]:((~r4_lattice3(X79,X80,X81)|(~m1_subset_1(X82,u1_struct_0(X79))|(~r2_hidden(X82,X81)|r1_lattices(X79,X82,X80)))|~m1_subset_1(X80,u1_struct_0(X79))|(v2_struct_0(X79)|~l3_lattices(X79)))&((m1_subset_1(esk13_3(X79,X80,X83),u1_struct_0(X79))|r4_lattice3(X79,X80,X83)|~m1_subset_1(X80,u1_struct_0(X79))|(v2_struct_0(X79)|~l3_lattices(X79)))&((r2_hidden(esk13_3(X79,X80,X83),X83)|r4_lattice3(X79,X80,X83)|~m1_subset_1(X80,u1_struct_0(X79))|(v2_struct_0(X79)|~l3_lattices(X79)))&(~r1_lattices(X79,esk13_3(X79,X80,X83),X80)|r4_lattice3(X79,X80,X83)|~m1_subset_1(X80,u1_struct_0(X79))|(v2_struct_0(X79)|~l3_lattices(X79)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_lattice3])])])])])])])).
cnf(c_0_119, negated_conjecture, (k15_lattice3(esk1_0,k1_xboole_0)=k15_lattice3(esk1_0,X1)|~r4_lattice3(esk1_0,k15_lattice3(esk1_0,k1_xboole_0),X1)|~r4_lattice3(esk1_0,k15_lattice3(esk1_0,X1),k1_xboole_0)|~m1_subset_1(k15_lattice3(esk1_0,k1_xboole_0),u1_struct_0(esk1_0))|~m1_subset_1(k15_lattice3(esk1_0,X1),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_77]), c_0_78]), c_0_79])]), c_0_80])).
cnf(c_0_120, plain, (r4_lattice3(X1,X2,a_2_3_lattice3(X1,X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_121, plain, (r2_hidden(esk13_3(X1,X2,X3),X3)|r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_122, negated_conjecture, (k15_lattice3(esk1_0,a_2_3_lattice3(esk1_0,k15_lattice3(esk1_0,k1_xboole_0)))=k15_lattice3(esk1_0,k1_xboole_0)|~r4_lattice3(esk1_0,k15_lattice3(esk1_0,a_2_3_lattice3(esk1_0,k15_lattice3(esk1_0,k1_xboole_0))),k1_xboole_0)|~m1_subset_1(k15_lattice3(esk1_0,a_2_3_lattice3(esk1_0,k15_lattice3(esk1_0,k1_xboole_0))),u1_struct_0(esk1_0))|~m1_subset_1(k15_lattice3(esk1_0,k1_xboole_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120]), c_0_77]), c_0_78]), c_0_79])]), c_0_80])).
cnf(c_0_123, plain, (r4_lattice3(X1,X2,k1_xboole_0)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_56, c_0_121])).
cnf(c_0_124, plain, (r1_lattices(X1,X4,X2)|v2_struct_0(X1)|~r4_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_125, negated_conjecture, (k15_lattice3(esk1_0,a_2_3_lattice3(esk1_0,k15_lattice3(esk1_0,k1_xboole_0)))=k15_lattice3(esk1_0,k1_xboole_0)|~m1_subset_1(k15_lattice3(esk1_0,a_2_3_lattice3(esk1_0,k15_lattice3(esk1_0,k1_xboole_0))),u1_struct_0(esk1_0))|~m1_subset_1(k15_lattice3(esk1_0,k1_xboole_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_77])]), c_0_80])).
fof(c_0_126, plain, ![X59, X60, X61, X63]:(((m1_subset_1(esk10_3(X59,X60,X61),u1_struct_0(X60))|~r2_hidden(X59,a_2_3_lattice3(X60,X61))|(v2_struct_0(X60)|(~v10_lattices(X60)|(~v4_lattice3(X60)|~l3_lattices(X60)))|~m1_subset_1(X61,u1_struct_0(X60))))&((X59=esk10_3(X59,X60,X61)|~r2_hidden(X59,a_2_3_lattice3(X60,X61))|(v2_struct_0(X60)|(~v10_lattices(X60)|(~v4_lattice3(X60)|~l3_lattices(X60)))|~m1_subset_1(X61,u1_struct_0(X60))))&(r3_lattices(X60,esk10_3(X59,X60,X61),X61)|~r2_hidden(X59,a_2_3_lattice3(X60,X61))|(v2_struct_0(X60)|(~v10_lattices(X60)|(~v4_lattice3(X60)|~l3_lattices(X60)))|~m1_subset_1(X61,u1_struct_0(X60))))))&(~m1_subset_1(X63,u1_struct_0(X60))|(X59!=X63|~r3_lattices(X60,X63,X61))|r2_hidden(X59,a_2_3_lattice3(X60,X61))|(v2_struct_0(X60)|(~v10_lattices(X60)|(~v4_lattice3(X60)|~l3_lattices(X60)))|~m1_subset_1(X61,u1_struct_0(X60))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_3_lattice3])])])])])])).
cnf(c_0_127, plain, (r1_lattices(X1,X2,k15_lattice3(X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_116]), c_0_95])).
cnf(c_0_128, negated_conjecture, (k15_lattice3(esk1_0,a_2_3_lattice3(esk1_0,k15_lattice3(esk1_0,k1_xboole_0)))=k15_lattice3(esk1_0,k1_xboole_0)|~m1_subset_1(k15_lattice3(esk1_0,k1_xboole_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_95]), c_0_77])]), c_0_80])).
cnf(c_0_129, plain, (r2_hidden(X3,a_2_3_lattice3(X2,X4))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~r3_lattices(X2,X1,X4)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_126])).
cnf(c_0_130, negated_conjecture, (r1_lattices(esk1_0,X1,k15_lattice3(esk1_0,k1_xboole_0))|~r2_hidden(X1,a_2_3_lattice3(esk1_0,k15_lattice3(esk1_0,k1_xboole_0)))|~m1_subset_1(k15_lattice3(esk1_0,k1_xboole_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128]), c_0_77]), c_0_78]), c_0_79])]), c_0_80])).
cnf(c_0_131, plain, (r2_hidden(X1,a_2_3_lattice3(X2,X3))|v2_struct_0(X2)|~r3_lattices(X2,X1,X3)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)), inference(er,[status(thm)],[c_0_129])).
fof(c_0_132, plain, ![X75, X76, X77]:((r3_lattices(X75,k15_lattice3(X75,X76),k15_lattice3(X75,X77))|~r1_tarski(X76,X77)|(v2_struct_0(X75)|(~v10_lattices(X75)|(~v4_lattice3(X75)|~l3_lattices(X75)))))&(r3_lattices(X75,k16_lattice3(X75,X77),k16_lattice3(X75,X76))|~r1_tarski(X76,X77)|(v2_struct_0(X75)|(~v10_lattices(X75)|(~v4_lattice3(X75)|~l3_lattices(X75)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t46_lattice3])])])])])).
cnf(c_0_133, negated_conjecture, (r1_lattices(esk1_0,X1,k15_lattice3(esk1_0,k1_xboole_0))|~r3_lattices(esk1_0,X1,k15_lattice3(esk1_0,k1_xboole_0))|~m1_subset_1(k15_lattice3(esk1_0,k1_xboole_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_131]), c_0_77]), c_0_78]), c_0_79])]), c_0_80])).
cnf(c_0_134, plain, (r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3))|v2_struct_0(X1)|~r1_tarski(X2,X3)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_135, negated_conjecture, (r1_lattices(esk1_0,k15_lattice3(esk1_0,X1),k15_lattice3(esk1_0,k1_xboole_0))|~r1_tarski(X1,k1_xboole_0)|~m1_subset_1(k15_lattice3(esk1_0,k1_xboole_0),u1_struct_0(esk1_0))|~m1_subset_1(k15_lattice3(esk1_0,X1),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134]), c_0_77]), c_0_78]), c_0_79])]), c_0_80])).
cnf(c_0_136, negated_conjecture, (k15_lattice3(esk1_0,X1)=k15_lattice3(esk1_0,k1_xboole_0)|~r4_lattice3(esk1_0,k15_lattice3(esk1_0,X1),k1_xboole_0)|~r1_tarski(X1,k1_xboole_0)|~m1_subset_1(k15_lattice3(esk1_0,k1_xboole_0),u1_struct_0(esk1_0))|~m1_subset_1(k15_lattice3(esk1_0,X1),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_135]), c_0_77]), c_0_78]), c_0_79])]), c_0_80])).
cnf(c_0_137, negated_conjecture, (k15_lattice3(esk1_0,X1)=k15_lattice3(esk1_0,k1_xboole_0)|~r1_tarski(X1,k1_xboole_0)|~m1_subset_1(k15_lattice3(esk1_0,k1_xboole_0),u1_struct_0(esk1_0))|~m1_subset_1(k15_lattice3(esk1_0,X1),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_123]), c_0_77])]), c_0_80])).
cnf(c_0_138, negated_conjecture, (k15_lattice3(esk1_0,X1)=k15_lattice3(esk1_0,k1_xboole_0)|~r1_tarski(X1,k1_xboole_0)|~m1_subset_1(k15_lattice3(esk1_0,X1),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_95]), c_0_77])]), c_0_80])).
fof(c_0_139, plain, ![X17, X18, X19]:((~v9_lattices(X17)|(~m1_subset_1(X18,u1_struct_0(X17))|(~m1_subset_1(X19,u1_struct_0(X17))|k2_lattices(X17,X18,k1_lattices(X17,X18,X19))=X18))|(v2_struct_0(X17)|~l3_lattices(X17)))&((m1_subset_1(esk3_1(X17),u1_struct_0(X17))|v9_lattices(X17)|(v2_struct_0(X17)|~l3_lattices(X17)))&((m1_subset_1(esk4_1(X17),u1_struct_0(X17))|v9_lattices(X17)|(v2_struct_0(X17)|~l3_lattices(X17)))&(k2_lattices(X17,esk3_1(X17),k1_lattices(X17,esk3_1(X17),esk4_1(X17)))!=esk3_1(X17)|v9_lattices(X17)|(v2_struct_0(X17)|~l3_lattices(X17)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d9_lattices])])])])])])).
fof(c_0_140, plain, ![X46, X47, X48]:((~r1_lattices(X46,X47,X48)|k1_lattices(X46,X47,X48)=X48|~m1_subset_1(X48,u1_struct_0(X46))|~m1_subset_1(X47,u1_struct_0(X46))|(v2_struct_0(X46)|~l2_lattices(X46)))&(k1_lattices(X46,X47,X48)!=X48|r1_lattices(X46,X47,X48)|~m1_subset_1(X48,u1_struct_0(X46))|~m1_subset_1(X47,u1_struct_0(X46))|(v2_struct_0(X46)|~l2_lattices(X46)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d3_lattices])])])])])).
cnf(c_0_141, plain, (r4_lattice3(X1,X2,a_2_5_lattice3(X3,k1_xboole_0))|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X3)|~l3_lattices(X1)|~v4_lattice3(X3)|~v10_lattices(X3)), inference(spm,[status(thm)],[c_0_61, c_0_121])).
cnf(c_0_142, negated_conjecture, (k15_lattice3(esk1_0,X1)=k15_lattice3(esk1_0,k1_xboole_0)|~r1_tarski(X1,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_95]), c_0_77])]), c_0_80])).
fof(c_0_143, plain, ![X85, X86, X87]:(((k2_lattices(X85,X86,X87)=X86|~m1_subset_1(X87,u1_struct_0(X85))|X86!=k5_lattices(X85)|~m1_subset_1(X86,u1_struct_0(X85))|~v13_lattices(X85)|(v2_struct_0(X85)|~l1_lattices(X85)))&(k2_lattices(X85,X87,X86)=X86|~m1_subset_1(X87,u1_struct_0(X85))|X86!=k5_lattices(X85)|~m1_subset_1(X86,u1_struct_0(X85))|~v13_lattices(X85)|(v2_struct_0(X85)|~l1_lattices(X85))))&((m1_subset_1(esk14_2(X85,X86),u1_struct_0(X85))|X86=k5_lattices(X85)|~m1_subset_1(X86,u1_struct_0(X85))|~v13_lattices(X85)|(v2_struct_0(X85)|~l1_lattices(X85)))&(k2_lattices(X85,X86,esk14_2(X85,X86))!=X86|k2_lattices(X85,esk14_2(X85,X86),X86)!=X86|X86=k5_lattices(X85)|~m1_subset_1(X86,u1_struct_0(X85))|~v13_lattices(X85)|(v2_struct_0(X85)|~l1_lattices(X85))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattices])])])])])])).
fof(c_0_144, plain, ![X7]:(v2_struct_0(X7)|~l1_lattices(X7)|m1_subset_1(k5_lattices(X7),u1_struct_0(X7))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_lattices])])])).
cnf(c_0_145, plain, (k2_lattices(X1,X2,k1_lattices(X1,X2,X3))=X2|v2_struct_0(X1)|~v9_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_139])).
cnf(c_0_146, plain, (k1_lattices(X1,X2,X3)=X3|v2_struct_0(X1)|~r1_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_140])).
cnf(c_0_147, plain, (r1_lattices(X1,k15_lattice3(X1,k1_xboole_0),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_110, c_0_141])).
cnf(c_0_148, negated_conjecture, (m1_subset_1(k15_lattice3(esk1_0,k1_xboole_0),u1_struct_0(esk1_0))|~r1_tarski(X1,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_142]), c_0_77])]), c_0_80])).
cnf(c_0_149, plain, (k2_lattices(X1,X2,X3)=X3|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|X3!=k5_lattices(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v13_lattices(X1)|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_143])).
cnf(c_0_150, plain, (v2_struct_0(X1)|m1_subset_1(k5_lattices(X1),u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_144])).
cnf(c_0_151, plain, (k2_lattices(X1,X2,X3)=X2|v2_struct_0(X1)|~v9_lattices(X1)|~r1_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146]), c_0_109])).
cnf(c_0_152, negated_conjecture, (r1_lattices(esk1_0,k15_lattice3(esk1_0,k1_xboole_0),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_142]), c_0_77]), c_0_78]), c_0_79]), c_0_39])]), c_0_80])).
cnf(c_0_153, negated_conjecture, (m1_subset_1(k15_lattice3(esk1_0,k1_xboole_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_148, c_0_39])).
cnf(c_0_154, negated_conjecture, (v2_struct_0(esk1_0)|~v10_lattices(esk1_0)|~v13_lattices(esk1_0)|~l3_lattices(esk1_0)|k5_lattices(esk1_0)!=k15_lattice3(esk1_0,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_73])).
fof(c_0_155, plain, ![X64, X65, X66]:((~v6_lattices(X64)|(~m1_subset_1(X65,u1_struct_0(X64))|(~m1_subset_1(X66,u1_struct_0(X64))|k2_lattices(X64,X65,X66)=k2_lattices(X64,X66,X65)))|(v2_struct_0(X64)|~l1_lattices(X64)))&((m1_subset_1(esk11_1(X64),u1_struct_0(X64))|v6_lattices(X64)|(v2_struct_0(X64)|~l1_lattices(X64)))&((m1_subset_1(esk12_1(X64),u1_struct_0(X64))|v6_lattices(X64)|(v2_struct_0(X64)|~l1_lattices(X64)))&(k2_lattices(X64,esk11_1(X64),esk12_1(X64))!=k2_lattices(X64,esk12_1(X64),esk11_1(X64))|v6_lattices(X64)|(v2_struct_0(X64)|~l1_lattices(X64)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_lattices])])])])])])).
cnf(c_0_156, plain, (k2_lattices(X1,X2,k5_lattices(X1))=k5_lattices(X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lattices(X1)|~v13_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_149]), c_0_150])).
cnf(c_0_157, negated_conjecture, (k2_lattices(esk1_0,k15_lattice3(esk1_0,k1_xboole_0),X1)=k15_lattice3(esk1_0,k1_xboole_0)|~v9_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152]), c_0_153]), c_0_77])]), c_0_80])).
cnf(c_0_158, negated_conjecture, (k5_lattices(esk1_0)!=k15_lattice3(esk1_0,k1_xboole_0)|~v13_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_154, c_0_79]), c_0_77])]), c_0_80])).
cnf(c_0_159, plain, (k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2)|v2_struct_0(X1)|~v6_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_155])).
cnf(c_0_160, plain, (v6_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_161, plain, (l1_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_162, negated_conjecture, (~v9_lattices(esk1_0)|~m1_subset_1(k5_lattices(esk1_0),u1_struct_0(esk1_0))|~l1_lattices(esk1_0)|~v13_lattices(esk1_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_157]), c_0_153])]), c_0_80]), c_0_158])).
fof(c_0_163, plain, ![X31, X33, X34]:(((m1_subset_1(esk5_1(X31),u1_struct_0(X31))|~v13_lattices(X31)|(v2_struct_0(X31)|~l1_lattices(X31)))&((k2_lattices(X31,esk5_1(X31),X33)=esk5_1(X31)|~m1_subset_1(X33,u1_struct_0(X31))|~v13_lattices(X31)|(v2_struct_0(X31)|~l1_lattices(X31)))&(k2_lattices(X31,X33,esk5_1(X31))=esk5_1(X31)|~m1_subset_1(X33,u1_struct_0(X31))|~v13_lattices(X31)|(v2_struct_0(X31)|~l1_lattices(X31)))))&((m1_subset_1(esk6_2(X31,X34),u1_struct_0(X31))|~m1_subset_1(X34,u1_struct_0(X31))|v13_lattices(X31)|(v2_struct_0(X31)|~l1_lattices(X31)))&(k2_lattices(X31,X34,esk6_2(X31,X34))!=X34|k2_lattices(X31,esk6_2(X31,X34),X34)!=X34|~m1_subset_1(X34,u1_struct_0(X31))|v13_lattices(X31)|(v2_struct_0(X31)|~l1_lattices(X31))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_lattices])])])])])])).
cnf(c_0_164, plain, (k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_160]), c_0_161])).
cnf(c_0_165, negated_conjecture, (~v9_lattices(esk1_0)|~l1_lattices(esk1_0)|~v13_lattices(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_150]), c_0_80])).
cnf(c_0_166, plain, (v9_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_167, plain, (v13_lattices(X1)|v2_struct_0(X1)|k2_lattices(X1,X2,esk6_2(X1,X2))!=X2|k2_lattices(X1,esk6_2(X1,X2),X2)!=X2|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_163])).
cnf(c_0_168, negated_conjecture, (k2_lattices(esk1_0,X1,k15_lattice3(esk1_0,k1_xboole_0))=k15_lattice3(esk1_0,k1_xboole_0)|~v9_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_157]), c_0_153]), c_0_77]), c_0_79])]), c_0_80])).
cnf(c_0_169, negated_conjecture, (~l1_lattices(esk1_0)|~v13_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_166]), c_0_77]), c_0_79])]), c_0_80])).
cnf(c_0_170, negated_conjecture, (~v9_lattices(esk1_0)|~m1_subset_1(esk6_2(esk1_0,k15_lattice3(esk1_0,k1_xboole_0)),u1_struct_0(esk1_0))|~l1_lattices(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_168]), c_0_153])]), c_0_80]), c_0_157]), c_0_169])).
cnf(c_0_171, plain, (m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1))|v13_lattices(X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_163])).
cnf(c_0_172, negated_conjecture, (~v9_lattices(esk1_0)|~l1_lattices(esk1_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_171]), c_0_153])]), c_0_80]), c_0_169])).
cnf(c_0_173, negated_conjecture, (~l1_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_166]), c_0_77]), c_0_79])]), c_0_80])).
cnf(c_0_174, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_161]), c_0_77])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 175
# Proof object clause steps            : 118
# Proof object formula steps           : 57
# Proof object conjectures             : 41
# Proof object clause conjectures      : 38
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 38
# Proof object initial formulas used   : 28
# Proof object generating inferences   : 73
# Proof object simplifying inferences  : 128
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 28
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 77
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 76
# Processed clauses                    : 2076
# ...of these trivial                  : 11
# ...subsumed                          : 1471
# ...remaining for further processing  : 594
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 31
# Backward-rewritten                   : 91
# Generated clauses                    : 5957
# ...of the previous two non-trivial   : 5495
# Contextual simplify-reflections      : 31
# Paramodulations                      : 5948
# Factorizations                       : 1
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 389
#    Positive orientable unit clauses  : 18
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 367
# Current number of unprocessed clauses: 2816
# ...number of literals in the above   : 16194
# Current number of archived formulas  : 0
# Current number of archived clauses   : 197
# Clause-clause subsumption calls (NU) : 71710
# Rec. Clause-clause subsumption calls : 19501
# Non-unit clause-clause subsumptions  : 1372
# Unit Clause-clause subsumption calls : 513
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 21
# BW rewrite match successes           : 7
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 114725

# -------------------------------------------------
# User time                : 0.095 s
# System time              : 0.000 s
# Total time               : 0.095 s
# Maximum resident set size: 31508 pages
