# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t23_relat_1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>(X3=k1_tarski(k4_tarski(X1,X2))=>(k1_relat_1(X3)=k1_tarski(X1)&k2_relat_1(X3)=k1_tarski(X2)))), file('data/MPTP/proofs/7ad3e2456f67432bb3848c5bf08a831a.E_input', t23_relat_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/7ad3e2456f67432bb3848c5bf08a831a.E_input', t69_enumset1)).
fof(t35_zfmisc_1, axiom, ![X1, X2]:k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2))=k1_tarski(k4_tarski(X1,X2)), file('data/MPTP/proofs/7ad3e2456f67432bb3848c5bf08a831a.E_input', t35_zfmisc_1)).
fof(t129_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,k1_tarski(X4)))<=>(r2_hidden(X1,X3)&X2=X4)), file('data/MPTP/proofs/7ad3e2456f67432bb3848c5bf08a831a.E_input', t129_zfmisc_1)).
fof(t128_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k1_tarski(X3),X4))<=>(X1=X3&r2_hidden(X2,X4))), file('data/MPTP/proofs/7ad3e2456f67432bb3848c5bf08a831a.E_input', t128_zfmisc_1)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/7ad3e2456f67432bb3848c5bf08a831a.E_input', d4_relat_1)).
fof(t34_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k1_tarski(X3),k1_tarski(X4)))<=>(X1=X3&X2=X4)), file('data/MPTP/proofs/7ad3e2456f67432bb3848c5bf08a831a.E_input', t34_zfmisc_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/7ad3e2456f67432bb3848c5bf08a831a.E_input', d5_relat_1)).
fof(c_0_8, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>(X3=k1_tarski(k4_tarski(X1,X2))=>(k1_relat_1(X3)=k1_tarski(X1)&k2_relat_1(X3)=k1_tarski(X2))))), inference(assume_negation,[status(cth)],[t23_relat_1])).
fof(c_0_9, negated_conjecture, (v1_relat_1(esk3_0)&(esk3_0=k1_tarski(k4_tarski(esk1_0,esk2_0))&(k1_relat_1(esk3_0)!=k1_tarski(esk1_0)|k2_relat_1(esk3_0)!=k1_tarski(esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
fof(c_0_10, plain, ![X8]:k2_tarski(X8,X8)=k1_tarski(X8), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_11, plain, ![X9, X10]:k2_zfmisc_1(k1_tarski(X9),k1_tarski(X10))=k1_tarski(k4_tarski(X9,X10)), inference(variable_rename,[status(thm)],[t35_zfmisc_1])).
fof(c_0_12, plain, ![X11, X12, X13, X14]:(((r2_hidden(X11,X13)|~r2_hidden(k4_tarski(X11,X12),k2_zfmisc_1(X13,k1_tarski(X14))))&(X12=X14|~r2_hidden(k4_tarski(X11,X12),k2_zfmisc_1(X13,k1_tarski(X14)))))&(~r2_hidden(X11,X13)|X12!=X14|r2_hidden(k4_tarski(X11,X12),k2_zfmisc_1(X13,k1_tarski(X14))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t129_zfmisc_1])])])).
cnf(c_0_13, negated_conjecture, (esk3_0=k1_tarski(k4_tarski(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, plain, (k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2))=k1_tarski(k4_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,k1_tarski(X4)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, negated_conjecture, (esk3_0=k2_tarski(k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_18, plain, (k2_zfmisc_1(k2_tarski(X1,X1),k2_tarski(X2,X2))=k2_tarski(k4_tarski(X1,X2),k4_tarski(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_14]), c_0_14]), c_0_14])).
fof(c_0_19, plain, ![X15, X16, X17, X18]:(((X15=X17|~r2_hidden(k4_tarski(X15,X16),k2_zfmisc_1(k1_tarski(X17),X18)))&(r2_hidden(X16,X18)|~r2_hidden(k4_tarski(X15,X16),k2_zfmisc_1(k1_tarski(X17),X18))))&(X15!=X17|~r2_hidden(X16,X18)|r2_hidden(k4_tarski(X15,X16),k2_zfmisc_1(k1_tarski(X17),X18)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t128_zfmisc_1])])])).
cnf(c_0_20, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,k1_tarski(X4)))|~r2_hidden(X1,X2)|X3!=X4), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,k2_tarski(X4,X4)))), inference(rw,[status(thm)],[c_0_16, c_0_14])).
cnf(c_0_22, negated_conjecture, (k2_zfmisc_1(k2_tarski(esk1_0,esk1_0),k2_tarski(esk2_0,esk2_0))=esk3_0), inference(rw,[status(thm)],[c_0_17, c_0_18])).
fof(c_0_23, plain, ![X19, X20, X21, X23, X24, X25, X26, X28]:(((~r2_hidden(X21,X20)|r2_hidden(k4_tarski(X21,esk4_3(X19,X20,X21)),X19)|X20!=k1_relat_1(X19))&(~r2_hidden(k4_tarski(X23,X24),X19)|r2_hidden(X23,X20)|X20!=k1_relat_1(X19)))&((~r2_hidden(esk5_2(X25,X26),X26)|~r2_hidden(k4_tarski(esk5_2(X25,X26),X28),X25)|X26=k1_relat_1(X25))&(r2_hidden(esk5_2(X25,X26),X26)|r2_hidden(k4_tarski(esk5_2(X25,X26),esk6_2(X25,X26)),X25)|X26=k1_relat_1(X25)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
fof(c_0_24, plain, ![X30, X31, X32, X33]:(((X30=X32|~r2_hidden(k4_tarski(X30,X31),k2_zfmisc_1(k1_tarski(X32),k1_tarski(X33))))&(X31=X33|~r2_hidden(k4_tarski(X30,X31),k2_zfmisc_1(k1_tarski(X32),k1_tarski(X33)))))&(X30!=X32|X31!=X33|r2_hidden(k4_tarski(X30,X31),k2_zfmisc_1(k1_tarski(X32),k1_tarski(X33))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t34_zfmisc_1])])])).
cnf(c_0_25, plain, (X1=X2|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k1_tarski(X2),X4))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,k2_tarski(X4,X4)))|X3!=X4|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_20, c_0_14])).
cnf(c_0_27, negated_conjecture, (r2_hidden(X1,k2_tarski(esk1_0,esk1_0))|~r2_hidden(k4_tarski(X1,X2),esk3_0)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28, plain, (r2_hidden(esk5_2(X1,X2),X2)|r2_hidden(k4_tarski(esk5_2(X1,X2),esk6_2(X1,X2)),X1)|X2=k1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k1_tarski(X2),k1_tarski(X4)))|X1!=X2|X3!=X4), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, plain, (X1=X2|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,k1_tarski(X2)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_31, plain, (X1=X2|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k2_tarski(X2,X2),X4))), inference(rw,[status(thm)],[c_0_25, c_0_14])).
cnf(c_0_32, plain, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,k2_tarski(X2,X2)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_26])).
cnf(c_0_33, negated_conjecture, (X1=k1_relat_1(esk3_0)|r2_hidden(esk5_2(esk3_0,X1),k2_tarski(esk1_0,esk1_0))|r2_hidden(esk5_2(esk3_0,X1),X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k2_tarski(X2,X2),k2_tarski(X4,X4)))|X3!=X4|X1!=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_14]), c_0_14])).
cnf(c_0_35, plain, (X1=X2|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,k2_tarski(X2,X2)))), inference(rw,[status(thm)],[c_0_30, c_0_14])).
fof(c_0_36, plain, ![X34, X35, X36, X38, X39, X40, X41, X43]:(((~r2_hidden(X36,X35)|r2_hidden(k4_tarski(esk7_3(X34,X35,X36),X36),X34)|X35!=k2_relat_1(X34))&(~r2_hidden(k4_tarski(X39,X38),X34)|r2_hidden(X38,X35)|X35!=k2_relat_1(X34)))&((~r2_hidden(esk8_2(X40,X41),X41)|~r2_hidden(k4_tarski(X43,esk8_2(X40,X41)),X40)|X41=k2_relat_1(X40))&(r2_hidden(esk8_2(X40,X41),X41)|r2_hidden(k4_tarski(esk9_2(X40,X41),esk8_2(X40,X41)),X40)|X41=k2_relat_1(X40)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_37, plain, (X1=X2|~r2_hidden(X1,k2_tarski(X2,X2))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_38, negated_conjecture, (k1_relat_1(esk3_0)=k2_tarski(esk1_0,esk1_0)|r2_hidden(esk5_2(esk3_0,k2_tarski(esk1_0,esk1_0)),k2_tarski(esk1_0,esk1_0))), inference(ef,[status(thm)],[c_0_33])).
cnf(c_0_39, plain, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k2_tarski(X1,X1),k2_tarski(X2,X2)))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_34])])).
cnf(c_0_40, negated_conjecture, (k1_relat_1(esk3_0)!=k1_tarski(esk1_0)|k2_relat_1(esk3_0)!=k1_tarski(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_41, negated_conjecture, (X1=esk2_0|~r2_hidden(k4_tarski(X2,X1),esk3_0)), inference(spm,[status(thm)],[c_0_35, c_0_22])).
cnf(c_0_42, plain, (r2_hidden(esk8_2(X1,X2),X2)|r2_hidden(k4_tarski(esk9_2(X1,X2),esk8_2(X1,X2)),X1)|X2=k2_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_43, plain, (X2=k1_relat_1(X1)|~r2_hidden(esk5_2(X1,X2),X2)|~r2_hidden(k4_tarski(esk5_2(X1,X2),X3),X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_44, negated_conjecture, (esk5_2(esk3_0,k2_tarski(esk1_0,esk1_0))=esk1_0|k1_relat_1(esk3_0)=k2_tarski(esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_45, plain, (r2_hidden(X1,k2_tarski(X1,X1))), inference(spm,[status(thm)],[c_0_21, c_0_39])).
cnf(c_0_46, negated_conjecture, (k1_relat_1(esk3_0)!=k2_tarski(esk1_0,esk1_0)|k2_relat_1(esk3_0)!=k2_tarski(esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_14]), c_0_14])).
cnf(c_0_47, negated_conjecture, (esk8_2(esk3_0,X1)=esk2_0|X1=k2_relat_1(esk3_0)|r2_hidden(esk8_2(esk3_0,X1),X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_48, negated_conjecture, (k1_relat_1(esk3_0)=k2_tarski(esk1_0,esk1_0)|~r2_hidden(k4_tarski(esk1_0,X1),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])])).
cnf(c_0_49, negated_conjecture, (r2_hidden(k4_tarski(X1,esk2_0),esk3_0)|~r2_hidden(X1,k2_tarski(esk1_0,esk1_0))), inference(spm,[status(thm)],[c_0_32, c_0_22])).
cnf(c_0_50, negated_conjecture, (esk8_2(esk3_0,X1)=esk2_0|r2_hidden(esk8_2(esk3_0,X1),X1)|k1_relat_1(esk3_0)!=k2_tarski(esk1_0,esk1_0)|X1!=k2_tarski(esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_51, negated_conjecture, (k1_relat_1(esk3_0)=k2_tarski(esk1_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_45])])).
cnf(c_0_52, negated_conjecture, (esk8_2(esk3_0,X1)=esk2_0|r2_hidden(esk8_2(esk3_0,X1),X1)|X1!=k2_tarski(esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_51])])).
cnf(c_0_53, plain, (X2=k2_relat_1(X1)|~r2_hidden(esk8_2(X1,X2),X2)|~r2_hidden(k4_tarski(X3,esk8_2(X1,X2)),X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_54, negated_conjecture, (esk8_2(esk3_0,k2_tarski(esk2_0,esk2_0))=esk2_0), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_52]), c_0_37])).
cnf(c_0_55, negated_conjecture, (k2_relat_1(esk3_0)!=k2_tarski(esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_51])])).
cnf(c_0_56, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_39, c_0_22])).
cnf(c_0_57, negated_conjecture, (~r2_hidden(k4_tarski(X1,esk2_0),esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_45])]), c_0_55])).
cnf(c_0_58, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_56, c_0_57]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 59
# Proof object clause steps            : 42
# Proof object formula steps           : 17
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 29
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 21
# Processed clauses                    : 775
# ...of these trivial                  : 33
# ...subsumed                          : 168
# ...remaining for further processing  : 574
# Other redundant clauses eliminated   : 168
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 31
# Backward-rewritten                   : 193
# Generated clauses                    : 6470
# ...of the previous two non-trivial   : 6209
# Contextual simplify-reflections      : 6
# Paramodulations                      : 6195
# Factorizations                       : 9
# Equation resolutions                 : 266
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 327
#    Positive orientable unit clauses  : 98
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 226
# Current number of unprocessed clauses: 5457
# ...number of literals in the above   : 19238
# Current number of archived formulas  : 0
# Current number of archived clauses   : 245
# Clause-clause subsumption calls (NU) : 13141
# Rec. Clause-clause subsumption calls : 8308
# Non-unit clause-clause subsumptions  : 185
# Unit Clause-clause subsumption calls : 1819
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1294
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 116934

# -------------------------------------------------
# User time                : 0.098 s
# System time              : 0.000 s
# Total time               : 0.098 s
# Maximum resident set size: 26200 pages
