# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t38_subset_1, conjecture, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(r1_tarski(X2,k3_subset_1(X1,X2))<=>X2=k1_subset_1(X1))), file('data/MPTP/proofs/7b12710ecff7463a81e46f7162ccb637.E_input', t38_subset_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/7b12710ecff7463a81e46f7162ccb637.E_input', d2_subset_1)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/7b12710ecff7463a81e46f7162ccb637.E_input', fc1_subset_1)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/7b12710ecff7463a81e46f7162ccb637.E_input', d1_zfmisc_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/7b12710ecff7463a81e46f7162ccb637.E_input', t36_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/7b12710ecff7463a81e46f7162ccb637.E_input', t100_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/7b12710ecff7463a81e46f7162ccb637.E_input', t28_xboole_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/7b12710ecff7463a81e46f7162ccb637.E_input', l32_xboole_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/7b12710ecff7463a81e46f7162ccb637.E_input', d5_xboole_0)).
fof(d3_subset_1, axiom, ![X1]:k1_subset_1(X1)=k1_xboole_0, file('data/MPTP/proofs/7b12710ecff7463a81e46f7162ccb637.E_input', d3_subset_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/7b12710ecff7463a81e46f7162ccb637.E_input', commutativity_k3_xboole_0)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/7b12710ecff7463a81e46f7162ccb637.E_input', d4_xboole_0)).
fof(d5_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)), file('data/MPTP/proofs/7b12710ecff7463a81e46f7162ccb637.E_input', d5_subset_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/7b12710ecff7463a81e46f7162ccb637.E_input', t7_xboole_0)).
fof(c_0_14, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(r1_tarski(X2,k3_subset_1(X1,X2))<=>X2=k1_subset_1(X1)))), inference(assume_negation,[status(cth)],[t38_subset_1])).
fof(c_0_15, plain, ![X7, X8]:(((~m1_subset_1(X8,X7)|r2_hidden(X8,X7)|v1_xboole_0(X7))&(~r2_hidden(X8,X7)|m1_subset_1(X8,X7)|v1_xboole_0(X7)))&((~m1_subset_1(X8,X7)|v1_xboole_0(X8)|~v1_xboole_0(X7))&(~v1_xboole_0(X8)|m1_subset_1(X8,X7)|~v1_xboole_0(X7)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_16, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))&((~r1_tarski(esk2_0,k3_subset_1(esk1_0,esk2_0))|esk2_0!=k1_subset_1(esk1_0))&(r1_tarski(esk2_0,k3_subset_1(esk1_0,esk2_0))|esk2_0=k1_subset_1(esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
fof(c_0_17, plain, ![X9]:~v1_xboole_0(k1_zfmisc_1(X9)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
fof(c_0_18, plain, ![X10, X11, X12, X13, X14, X15]:(((~r2_hidden(X12,X11)|r1_tarski(X12,X10)|X11!=k1_zfmisc_1(X10))&(~r1_tarski(X13,X10)|r2_hidden(X13,X11)|X11!=k1_zfmisc_1(X10)))&((~r2_hidden(esk3_2(X14,X15),X15)|~r1_tarski(esk3_2(X14,X15),X14)|X15=k1_zfmisc_1(X14))&(r2_hidden(esk3_2(X14,X15),X15)|r1_tarski(esk3_2(X14,X15),X14)|X15=k1_zfmisc_1(X14)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
cnf(c_0_19, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, negated_conjecture, (r2_hidden(esk2_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])).
fof(c_0_24, plain, ![X17, X18]:r1_tarski(k4_xboole_0(X17,X18),X17), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_25, plain, ![X19, X20]:k4_xboole_0(X19,X20)=k5_xboole_0(X19,k3_xboole_0(X19,X20)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_26, plain, ![X21, X22]:(~r1_tarski(X21,X22)|k3_xboole_0(X21,X22)=X21), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_27, negated_conjecture, (r1_tarski(esk2_0,X1)|k1_zfmisc_1(esk1_0)!=k1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
fof(c_0_28, plain, ![X23, X24]:((k4_xboole_0(X23,X24)!=k1_xboole_0|r1_tarski(X23,X24))&(~r1_tarski(X23,X24)|k4_xboole_0(X23,X24)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
fof(c_0_29, plain, ![X25, X26, X27, X28, X29, X30, X31, X32]:((((r2_hidden(X28,X25)|~r2_hidden(X28,X27)|X27!=k4_xboole_0(X25,X26))&(~r2_hidden(X28,X26)|~r2_hidden(X28,X27)|X27!=k4_xboole_0(X25,X26)))&(~r2_hidden(X29,X25)|r2_hidden(X29,X26)|r2_hidden(X29,X27)|X27!=k4_xboole_0(X25,X26)))&((~r2_hidden(esk4_3(X30,X31,X32),X32)|(~r2_hidden(esk4_3(X30,X31,X32),X30)|r2_hidden(esk4_3(X30,X31,X32),X31))|X32=k4_xboole_0(X30,X31))&((r2_hidden(esk4_3(X30,X31,X32),X30)|r2_hidden(esk4_3(X30,X31,X32),X32)|X32=k4_xboole_0(X30,X31))&(~r2_hidden(esk4_3(X30,X31,X32),X31)|r2_hidden(esk4_3(X30,X31,X32),X32)|X32=k4_xboole_0(X30,X31))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_30, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, negated_conjecture, (r1_tarski(esk2_0,esk1_0)), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_34, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_35, plain, ![X34]:k1_subset_1(X34)=k1_xboole_0, inference(variable_rename,[status(thm)],[d3_subset_1])).
cnf(c_0_36, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_37, plain, ![X35, X36]:k3_xboole_0(X35,X36)=k3_xboole_0(X36,X35), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_38, plain, (r1_tarski(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X1)), inference(rw,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_39, negated_conjecture, (k3_xboole_0(esk2_0,esk1_0)=esk2_0), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_40, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=k1_xboole_0|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_34, c_0_31])).
cnf(c_0_41, negated_conjecture, (r1_tarski(esk2_0,k3_subset_1(esk1_0,esk2_0))|esk2_0=k1_subset_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_42, plain, (k1_subset_1(X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_43, plain, (X3!=k5_xboole_0(X4,k3_xboole_0(X4,X2))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_36, c_0_31])).
cnf(c_0_44, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_45, negated_conjecture, (r1_tarski(k5_xboole_0(esk2_0,esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_46, negated_conjecture, (k1_xboole_0=k5_xboole_0(esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_33]), c_0_39])).
fof(c_0_47, plain, ![X37, X38, X39, X40, X41, X42, X43, X44]:((((r2_hidden(X40,X37)|~r2_hidden(X40,X39)|X39!=k3_xboole_0(X37,X38))&(r2_hidden(X40,X38)|~r2_hidden(X40,X39)|X39!=k3_xboole_0(X37,X38)))&(~r2_hidden(X41,X37)|~r2_hidden(X41,X38)|r2_hidden(X41,X39)|X39!=k3_xboole_0(X37,X38)))&((~r2_hidden(esk5_3(X42,X43,X44),X44)|(~r2_hidden(esk5_3(X42,X43,X44),X42)|~r2_hidden(esk5_3(X42,X43,X44),X43))|X44=k3_xboole_0(X42,X43))&((r2_hidden(esk5_3(X42,X43,X44),X42)|r2_hidden(esk5_3(X42,X43,X44),X44)|X44=k3_xboole_0(X42,X43))&(r2_hidden(esk5_3(X42,X43,X44),X43)|r2_hidden(esk5_3(X42,X43,X44),X44)|X44=k3_xboole_0(X42,X43))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_48, negated_conjecture, (esk2_0=k1_xboole_0|r1_tarski(esk2_0,k3_subset_1(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_49, plain, (X1!=k5_xboole_0(X2,k3_xboole_0(X3,X2))|~r2_hidden(X4,X1)|~r2_hidden(X4,X3)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_50, negated_conjecture, (k3_xboole_0(esk2_0,k5_xboole_0(esk2_0,esk2_0))=k5_xboole_0(esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_45]), c_0_44])).
cnf(c_0_51, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=k5_xboole_0(esk2_0,esk2_0)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_40, c_0_46])).
cnf(c_0_52, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_53, plain, ![X46, X47]:(~m1_subset_1(X47,k1_zfmisc_1(X46))|k3_subset_1(X46,X47)=k4_xboole_0(X46,X47)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
cnf(c_0_54, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_55, negated_conjecture, (k3_xboole_0(esk2_0,k3_subset_1(esk1_0,esk2_0))=esk2_0|k1_xboole_0=esk2_0), inference(spm,[status(thm)],[c_0_32, c_0_48])).
cnf(c_0_56, negated_conjecture, (X1!=k5_xboole_0(k5_xboole_0(esk2_0,esk2_0),k5_xboole_0(esk2_0,esk2_0))|~r2_hidden(X2,esk2_0)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_57, negated_conjecture, (k5_xboole_0(k5_xboole_0(esk2_0,esk2_0),k5_xboole_0(esk2_0,esk2_0))=k5_xboole_0(esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_45]), c_0_44]), c_0_50])).
cnf(c_0_58, negated_conjecture, (r2_hidden(X1,esk2_0)|X2!=k5_xboole_0(esk2_0,esk2_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_52, c_0_50])).
cnf(c_0_59, plain, (~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_43])).
cnf(c_0_60, negated_conjecture, (~r1_tarski(esk2_0,k3_subset_1(esk1_0,esk2_0))|esk2_0!=k1_subset_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_61, plain, (k3_subset_1(X2,X1)=k4_xboole_0(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_62, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_54])).
cnf(c_0_63, negated_conjecture, (k3_xboole_0(esk2_0,k3_subset_1(esk1_0,esk2_0))=esk2_0|k5_xboole_0(esk2_0,esk2_0)=esk2_0), inference(rw,[status(thm)],[c_0_55, c_0_46])).
cnf(c_0_64, negated_conjecture, (X1!=k5_xboole_0(esk2_0,esk2_0)|~r2_hidden(X2,X1)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_57]), c_0_58])).
cnf(c_0_65, plain, (~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X3,X2)))|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_59, c_0_44])).
fof(c_0_66, plain, ![X48]:(X48=k1_xboole_0|r2_hidden(esk6_1(X48),X48)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_67, plain, (r1_tarski(k5_xboole_0(X1,k3_xboole_0(X2,X1)),X1)), inference(spm,[status(thm)],[c_0_38, c_0_44])).
cnf(c_0_68, negated_conjecture, (esk2_0!=k1_xboole_0|~r1_tarski(esk2_0,k3_subset_1(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_60, c_0_42])).
cnf(c_0_69, plain, (k5_xboole_0(X2,k3_xboole_0(X2,X1))=k3_subset_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_61, c_0_31])).
cnf(c_0_70, negated_conjecture, (r2_hidden(X1,k3_subset_1(esk1_0,esk2_0))|~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64])).
cnf(c_0_71, negated_conjecture, (~r2_hidden(X1,k5_xboole_0(esk1_0,esk2_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_65, c_0_39])).
cnf(c_0_72, plain, (X1=k1_xboole_0|r2_hidden(esk6_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_73, negated_conjecture, (r1_tarski(k5_xboole_0(esk1_0,esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_67, c_0_39])).
cnf(c_0_74, negated_conjecture, (k1_xboole_0!=esk2_0|~r1_tarski(esk2_0,k5_xboole_0(esk1_0,esk2_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_20])]), c_0_44]), c_0_39])).
cnf(c_0_75, negated_conjecture, (~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_69]), c_0_44]), c_0_39]), c_0_20])]), c_0_71])).
cnf(c_0_76, plain, (X1=k5_xboole_0(esk2_0,esk2_0)|r2_hidden(esk6_1(X1),X1)), inference(rw,[status(thm)],[c_0_72, c_0_46])).
cnf(c_0_77, negated_conjecture, (k3_xboole_0(esk1_0,k5_xboole_0(esk1_0,esk2_0))=k5_xboole_0(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_73]), c_0_44])).
cnf(c_0_78, negated_conjecture, (k5_xboole_0(esk2_0,esk2_0)!=esk2_0|~r1_tarski(esk2_0,k5_xboole_0(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_74, c_0_46])).
cnf(c_0_79, negated_conjecture, (k5_xboole_0(esk2_0,esk2_0)=esk2_0), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_80, negated_conjecture, (k5_xboole_0(k5_xboole_0(esk1_0,esk2_0),k5_xboole_0(esk1_0,esk2_0))=k5_xboole_0(esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_73]), c_0_44]), c_0_77])).
cnf(c_0_81, negated_conjecture, (~r1_tarski(esk2_0,k5_xboole_0(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_79])])).
cnf(c_0_82, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_77]), c_0_80]), c_0_79]), c_0_81]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 83
# Proof object clause steps            : 54
# Proof object formula steps           : 29
# Proof object conjectures             : 32
# Proof object clause conjectures      : 29
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 35
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 31
# Processed clauses                    : 453
# ...of these trivial                  : 20
# ...subsumed                          : 174
# ...remaining for further processing  : 258
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 106
# Generated clauses                    : 815
# ...of the previous two non-trivial   : 812
# Contextual simplify-reflections      : 5
# Paramodulations                      : 757
# Factorizations                       : 6
# Equation resolutions                 : 52
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 117
#    Positive orientable unit clauses  : 21
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 88
# Current number of unprocessed clauses: 372
# ...number of literals in the above   : 1053
# Current number of archived formulas  : 0
# Current number of archived clauses   : 143
# Clause-clause subsumption calls (NU) : 3359
# Rec. Clause-clause subsumption calls : 2712
# Non-unit clause-clause subsumptions  : 111
# Unit Clause-clause subsumption calls : 337
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 15
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 12245

# -------------------------------------------------
# User time                : 0.017 s
# System time              : 0.003 s
# Total time               : 0.020 s
# Maximum resident set size: 28664 pages
