# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t24_yellow_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1)), file('data/MPTP/proofs/7c17d69392a4414d9eb06ee53d19a422.E_input', t24_yellow_1)).
fof(t56_setfam_1, axiom, ![X1, X2, X3]:((r1_tarski(k3_tarski(X1),X2)&r2_hidden(X3,X1))=>r1_tarski(X3,X2)), file('data/MPTP/proofs/7c17d69392a4414d9eb06ee53d19a422.E_input', t56_setfam_1)).
fof(t95_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>r1_tarski(k3_tarski(X1),k3_tarski(X2))), file('data/MPTP/proofs/7c17d69392a4414d9eb06ee53d19a422.E_input', t95_zfmisc_1)).
fof(d1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1))))))))), file('data/MPTP/proofs/7c17d69392a4414d9eb06ee53d19a422.E_input', d1_pre_topc)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/7c17d69392a4414d9eb06ee53d19a422.E_input', t3_subset)).
fof(d12_setfam_1, axiom, ![X1, X2]:(m1_setfam_1(X2,X1)<=>r1_tarski(X1,k3_tarski(X2))), file('data/MPTP/proofs/7c17d69392a4414d9eb06ee53d19a422.E_input', d12_setfam_1)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/7c17d69392a4414d9eb06ee53d19a422.E_input', t5_subset)).
fof(t60_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>(m1_setfam_1(X2,X1)<=>k5_setfam_1(X1,X2)=X1)), file('data/MPTP/proofs/7c17d69392a4414d9eb06ee53d19a422.E_input', t60_setfam_1)).
fof(t99_zfmisc_1, axiom, ![X1]:k3_tarski(k1_zfmisc_1(X1))=X1, file('data/MPTP/proofs/7c17d69392a4414d9eb06ee53d19a422.E_input', t99_zfmisc_1)).
fof(redefinition_k5_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>k5_setfam_1(X1,X2)=k3_tarski(X2)), file('data/MPTP/proofs/7c17d69392a4414d9eb06ee53d19a422.E_input', redefinition_k5_setfam_1)).
fof(t94_zfmisc_1, axiom, ![X1, X2]:(![X3]:(r2_hidden(X3,X1)=>r1_tarski(X3,X2))=>r1_tarski(k3_tarski(X1),X2)), file('data/MPTP/proofs/7c17d69392a4414d9eb06ee53d19a422.E_input', t94_zfmisc_1)).
fof(l49_zfmisc_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>r1_tarski(X1,k3_tarski(X2))), file('data/MPTP/proofs/7c17d69392a4414d9eb06ee53d19a422.E_input', l49_zfmisc_1)).
fof(dt_u1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/7c17d69392a4414d9eb06ee53d19a422.E_input', dt_u1_pre_topc)).
fof(t14_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>(r2_hidden(k3_tarski(X1),X1)=>k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1))), file('data/MPTP/proofs/7c17d69392a4414d9eb06ee53d19a422.E_input', t14_yellow_1)).
fof(c_0_14, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1))), inference(assume_negation,[status(cth)],[t24_yellow_1])).
fof(c_0_15, plain, ![X5, X6, X7]:(~r1_tarski(k3_tarski(X5),X6)|~r2_hidden(X7,X5)|r1_tarski(X7,X6)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t56_setfam_1])])).
fof(c_0_16, plain, ![X8, X9]:(~r1_tarski(X8,X9)|r1_tarski(k3_tarski(X8),k3_tarski(X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t95_zfmisc_1])])).
fof(c_0_17, plain, ![X10, X11, X12, X13]:(((r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|~v2_pre_topc(X10)|~l1_pre_topc(X10))&((~m1_subset_1(X11,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|(~r1_tarski(X11,u1_pre_topc(X10))|r2_hidden(k5_setfam_1(u1_struct_0(X10),X11),u1_pre_topc(X10)))|~v2_pre_topc(X10)|~l1_pre_topc(X10))&(~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X10)))|(~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X10)))|(~r2_hidden(X12,u1_pre_topc(X10))|~r2_hidden(X13,u1_pre_topc(X10))|r2_hidden(k9_subset_1(u1_struct_0(X10),X12,X13),u1_pre_topc(X10))))|~v2_pre_topc(X10)|~l1_pre_topc(X10))))&(((m1_subset_1(esk3_1(X10),k1_zfmisc_1(u1_struct_0(X10)))|m1_subset_1(esk2_1(X10),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&((m1_subset_1(esk4_1(X10),k1_zfmisc_1(u1_struct_0(X10)))|m1_subset_1(esk2_1(X10),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&(((r2_hidden(esk3_1(X10),u1_pre_topc(X10))|m1_subset_1(esk2_1(X10),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&(r2_hidden(esk4_1(X10),u1_pre_topc(X10))|m1_subset_1(esk2_1(X10),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10)))&(~r2_hidden(k9_subset_1(u1_struct_0(X10),esk3_1(X10),esk4_1(X10)),u1_pre_topc(X10))|m1_subset_1(esk2_1(X10),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10)))))&(((m1_subset_1(esk3_1(X10),k1_zfmisc_1(u1_struct_0(X10)))|r1_tarski(esk2_1(X10),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&((m1_subset_1(esk4_1(X10),k1_zfmisc_1(u1_struct_0(X10)))|r1_tarski(esk2_1(X10),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&(((r2_hidden(esk3_1(X10),u1_pre_topc(X10))|r1_tarski(esk2_1(X10),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&(r2_hidden(esk4_1(X10),u1_pre_topc(X10))|r1_tarski(esk2_1(X10),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10)))&(~r2_hidden(k9_subset_1(u1_struct_0(X10),esk3_1(X10),esk4_1(X10)),u1_pre_topc(X10))|r1_tarski(esk2_1(X10),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10)))))&((m1_subset_1(esk3_1(X10),k1_zfmisc_1(u1_struct_0(X10)))|~r2_hidden(k5_setfam_1(u1_struct_0(X10),esk2_1(X10)),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&((m1_subset_1(esk4_1(X10),k1_zfmisc_1(u1_struct_0(X10)))|~r2_hidden(k5_setfam_1(u1_struct_0(X10),esk2_1(X10)),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&(((r2_hidden(esk3_1(X10),u1_pre_topc(X10))|~r2_hidden(k5_setfam_1(u1_struct_0(X10),esk2_1(X10)),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10))&(r2_hidden(esk4_1(X10),u1_pre_topc(X10))|~r2_hidden(k5_setfam_1(u1_struct_0(X10),esk2_1(X10)),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10)))&(~r2_hidden(k9_subset_1(u1_struct_0(X10),esk3_1(X10),esk4_1(X10)),u1_pre_topc(X10))|~r2_hidden(k5_setfam_1(u1_struct_0(X10),esk2_1(X10)),u1_pre_topc(X10))|~r2_hidden(u1_struct_0(X10),u1_pre_topc(X10))|v2_pre_topc(X10)|~l1_pre_topc(X10)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])).
fof(c_0_18, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&k4_yellow_0(k2_yellow_1(u1_pre_topc(esk1_0)))!=u1_struct_0(esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])).
cnf(c_0_19, plain, (r1_tarski(X3,X2)|~r1_tarski(k3_tarski(X1),X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, plain, (r1_tarski(k3_tarski(X1),k3_tarski(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, plain, (r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_24, plain, ![X17, X18]:((~m1_subset_1(X17,k1_zfmisc_1(X18))|r1_tarski(X17,X18))&(~r1_tarski(X17,X18)|m1_subset_1(X17,k1_zfmisc_1(X18)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_25, plain, ![X19, X20]:((~m1_setfam_1(X20,X19)|r1_tarski(X19,k3_tarski(X20)))&(~r1_tarski(X19,k3_tarski(X20))|m1_setfam_1(X20,X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_setfam_1])])).
cnf(c_0_26, plain, (r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_27, negated_conjecture, (r2_hidden(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])])).
fof(c_0_28, plain, ![X21, X22, X23]:(~r2_hidden(X21,X22)|(~m1_subset_1(X22,k1_zfmisc_1(X23))|~v1_xboole_0(X23))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
cnf(c_0_29, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, plain, (r1_tarski(X2,k3_tarski(X1))|~m1_setfam_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_31, plain, ![X24, X25]:((~m1_setfam_1(X25,X24)|k5_setfam_1(X24,X25)=X24|~m1_subset_1(X25,k1_zfmisc_1(k1_zfmisc_1(X24))))&(k5_setfam_1(X24,X25)!=X24|m1_setfam_1(X25,X24)|~m1_subset_1(X25,k1_zfmisc_1(k1_zfmisc_1(X24))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t60_setfam_1])])])).
cnf(c_0_32, plain, (m1_setfam_1(X2,X1)|~r1_tarski(X1,k3_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, negated_conjecture, (r1_tarski(u1_struct_0(esk1_0),k3_tarski(X1))|~r1_tarski(u1_pre_topc(esk1_0),X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_34, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, plain, (m1_subset_1(X1,k1_zfmisc_1(k3_tarski(X2)))|~m1_setfam_1(X2,X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
fof(c_0_36, plain, ![X26]:k3_tarski(k1_zfmisc_1(X26))=X26, inference(variable_rename,[status(thm)],[t99_zfmisc_1])).
fof(c_0_37, plain, ![X27, X28]:(~m1_subset_1(X28,k1_zfmisc_1(k1_zfmisc_1(X27)))|k5_setfam_1(X27,X28)=k3_tarski(X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k5_setfam_1])])).
cnf(c_0_38, plain, (k5_setfam_1(X2,X1)=X2|~m1_setfam_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_39, negated_conjecture, (m1_setfam_1(X1,u1_struct_0(esk1_0))|~r1_tarski(u1_pre_topc(esk1_0),X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_40, plain, (~v1_xboole_0(k3_tarski(X1))|~m1_setfam_1(X1,X2)|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_41, plain, (k3_tarski(k1_zfmisc_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_42, plain, ![X29, X30]:((r2_hidden(esk5_2(X29,X30),X29)|r1_tarski(k3_tarski(X29),X30))&(~r1_tarski(esk5_2(X29,X30),X30)|r1_tarski(k3_tarski(X29),X30))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t94_zfmisc_1])])])])).
fof(c_0_43, plain, ![X32, X33]:(~r2_hidden(X32,X33)|r1_tarski(X32,k3_tarski(X33))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l49_zfmisc_1])])).
cnf(c_0_44, plain, (k5_setfam_1(X2,X1)=k3_tarski(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_45, negated_conjecture, (k5_setfam_1(u1_struct_0(esk1_0),X1)=u1_struct_0(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~r1_tarski(u1_pre_topc(esk1_0),X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
fof(c_0_46, plain, ![X34]:(~l1_pre_topc(X34)|m1_subset_1(u1_pre_topc(X34),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X34))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_pre_topc])])).
cnf(c_0_47, plain, (~v1_xboole_0(X1)|~m1_setfam_1(k1_zfmisc_1(X1),X2)|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_48, plain, (m1_setfam_1(k1_zfmisc_1(X1),X2)|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_32, c_0_41])).
cnf(c_0_49, plain, (r1_tarski(k3_tarski(X1),X2)|~r1_tarski(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_50, plain, (r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_51, plain, ![X35]:(v1_xboole_0(X35)|(~r2_hidden(k3_tarski(X35),X35)|k4_yellow_0(k2_yellow_1(X35))=k3_tarski(X35))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_yellow_1])])])).
cnf(c_0_52, negated_conjecture, (k3_tarski(X1)=u1_struct_0(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~r1_tarski(u1_pre_topc(esk1_0),X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_53, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_54, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X3)|~r1_tarski(X3,X1)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_55, plain, (r1_tarski(k3_tarski(X1),k3_tarski(X2))|~r2_hidden(esk5_2(X1,k3_tarski(X2)),X2)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_56, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(k3_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_57, plain, (v1_xboole_0(X1)|k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_58, negated_conjecture, (k3_tarski(u1_pre_topc(esk1_0))=u1_struct_0(esk1_0)|~r1_tarski(u1_pre_topc(esk1_0),u1_pre_topc(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_23])])).
cnf(c_0_59, negated_conjecture, (k4_yellow_0(k2_yellow_1(u1_pre_topc(esk1_0)))!=u1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_60, negated_conjecture, (~v1_xboole_0(X1)|~r1_tarski(u1_pre_topc(esk1_0),X1)), inference(spm,[status(thm)],[c_0_54, c_0_27])).
cnf(c_0_61, plain, (r1_tarski(k3_tarski(X1),k3_tarski(X1))), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_62, negated_conjecture, (~r1_tarski(u1_pre_topc(esk1_0),u1_pre_topc(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_27])]), c_0_59]), c_0_60])).
cnf(c_0_63, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_61, c_0_41])).
cnf(c_0_64, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_63])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 65
# Proof object clause steps            : 36
# Proof object formula steps           : 29
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 10
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 38
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 38
# Processed clauses                    : 223
# ...of these trivial                  : 0
# ...subsumed                          : 55
# ...remaining for further processing  : 168
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 4
# Generated clauses                    : 347
# ...of the previous two non-trivial   : 306
# Contextual simplify-reflections      : 1
# Paramodulations                      : 346
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 124
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 114
# Current number of unprocessed clauses: 158
# ...number of literals in the above   : 576
# Current number of archived formulas  : 0
# Current number of archived clauses   : 44
# Clause-clause subsumption calls (NU) : 2497
# Rec. Clause-clause subsumption calls : 1386
# Non-unit clause-clause subsumptions  : 52
# Unit Clause-clause subsumption calls : 64
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 39
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8570

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.003 s
# Total time               : 0.015 s
# Maximum resident set size: 28728 pages
