# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t22_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k13_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))))))), file('data/MPTP/proofs/7c50220d084347549d69dd8ee1253012.E_input', t22_yellow_0)).
fof(dt_k13_lattice3, axiom, ![X1, X2, X3]:(((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k13_lattice3(X1,X2,X3),u1_struct_0(X1))), file('data/MPTP/proofs/7c50220d084347549d69dd8ee1253012.E_input', dt_k13_lattice3)).
fof(d1_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v1_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X3,X2)&(r2_hidden(X4,X2)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>~((r2_hidden(X5,X2)&(r1_orders_2(X1,X3,X5)&r1_orders_2(X1,X4,X5))))))))))))), file('data/MPTP/proofs/7c50220d084347549d69dd8ee1253012.E_input', d1_waybel_0)).
fof(t40_waybel_0, conjecture, ![X1]:((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))=>![X2]:((v12_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r2_hidden(X4,X2))=>r2_hidden(k13_lattice3(X1,X3,X4),X2))))))), file('data/MPTP/proofs/7c50220d084347549d69dd8ee1253012.E_input', t40_waybel_0)).
fof(d19_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v12_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r1_orders_2(X1,X4,X3))=>r2_hidden(X4,X2))))))), file('data/MPTP/proofs/7c50220d084347549d69dd8ee1253012.E_input', d19_waybel_0)).
fof(c_0_5, plain, ![X29, X30, X31, X32, X33]:(((r1_orders_2(X29,X30,X32)|X32!=k13_lattice3(X29,X30,X31)|~m1_subset_1(X32,u1_struct_0(X29))|~m1_subset_1(X31,u1_struct_0(X29))|~m1_subset_1(X30,u1_struct_0(X29))|(~v5_orders_2(X29)|(~v1_lattice3(X29)|~l1_orders_2(X29))))&((r1_orders_2(X29,X31,X32)|X32!=k13_lattice3(X29,X30,X31)|~m1_subset_1(X32,u1_struct_0(X29))|~m1_subset_1(X31,u1_struct_0(X29))|~m1_subset_1(X30,u1_struct_0(X29))|(~v5_orders_2(X29)|(~v1_lattice3(X29)|~l1_orders_2(X29))))&(~m1_subset_1(X33,u1_struct_0(X29))|(~r1_orders_2(X29,X30,X33)|~r1_orders_2(X29,X31,X33)|r1_orders_2(X29,X32,X33))|X32!=k13_lattice3(X29,X30,X31)|~m1_subset_1(X32,u1_struct_0(X29))|~m1_subset_1(X31,u1_struct_0(X29))|~m1_subset_1(X30,u1_struct_0(X29))|(~v5_orders_2(X29)|(~v1_lattice3(X29)|~l1_orders_2(X29))))))&((m1_subset_1(esk10_4(X29,X30,X31,X32),u1_struct_0(X29))|~r1_orders_2(X29,X31,X32)|~r1_orders_2(X29,X30,X32)|X32=k13_lattice3(X29,X30,X31)|~m1_subset_1(X32,u1_struct_0(X29))|~m1_subset_1(X31,u1_struct_0(X29))|~m1_subset_1(X30,u1_struct_0(X29))|(~v5_orders_2(X29)|(~v1_lattice3(X29)|~l1_orders_2(X29))))&(((r1_orders_2(X29,X30,esk10_4(X29,X30,X31,X32))|~r1_orders_2(X29,X31,X32)|~r1_orders_2(X29,X30,X32)|X32=k13_lattice3(X29,X30,X31)|~m1_subset_1(X32,u1_struct_0(X29))|~m1_subset_1(X31,u1_struct_0(X29))|~m1_subset_1(X30,u1_struct_0(X29))|(~v5_orders_2(X29)|(~v1_lattice3(X29)|~l1_orders_2(X29))))&(r1_orders_2(X29,X31,esk10_4(X29,X30,X31,X32))|~r1_orders_2(X29,X31,X32)|~r1_orders_2(X29,X30,X32)|X32=k13_lattice3(X29,X30,X31)|~m1_subset_1(X32,u1_struct_0(X29))|~m1_subset_1(X31,u1_struct_0(X29))|~m1_subset_1(X30,u1_struct_0(X29))|(~v5_orders_2(X29)|(~v1_lattice3(X29)|~l1_orders_2(X29)))))&(~r1_orders_2(X29,X32,esk10_4(X29,X30,X31,X32))|~r1_orders_2(X29,X31,X32)|~r1_orders_2(X29,X30,X32)|X32=k13_lattice3(X29,X30,X31)|~m1_subset_1(X32,u1_struct_0(X29))|~m1_subset_1(X31,u1_struct_0(X29))|~m1_subset_1(X30,u1_struct_0(X29))|(~v5_orders_2(X29)|(~v1_lattice3(X29)|~l1_orders_2(X29))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t22_yellow_0])])])])])).
fof(c_0_6, plain, ![X26, X27, X28]:(~v5_orders_2(X26)|(~v1_lattice3(X26)|~l1_orders_2(X26))|(~m1_subset_1(X27,u1_struct_0(X26))|~m1_subset_1(X28,u1_struct_0(X26)))|m1_subset_1(k13_lattice3(X26,X27,X28),u1_struct_0(X26))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k13_lattice3])])).
fof(c_0_7, plain, ![X18, X19, X20, X21, X25]:(((m1_subset_1(esk7_4(X18,X19,X20,X21),u1_struct_0(X18))|~r2_hidden(X21,X19)|~r2_hidden(X20,X19)|~m1_subset_1(X21,u1_struct_0(X18))|~m1_subset_1(X20,u1_struct_0(X18))|~v1_waybel_0(X19,X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|~l1_orders_2(X18))&((r2_hidden(esk7_4(X18,X19,X20,X21),X19)|~r2_hidden(X21,X19)|~r2_hidden(X20,X19)|~m1_subset_1(X21,u1_struct_0(X18))|~m1_subset_1(X20,u1_struct_0(X18))|~v1_waybel_0(X19,X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|~l1_orders_2(X18))&((r1_orders_2(X18,X20,esk7_4(X18,X19,X20,X21))|~r2_hidden(X21,X19)|~r2_hidden(X20,X19)|~m1_subset_1(X21,u1_struct_0(X18))|~m1_subset_1(X20,u1_struct_0(X18))|~v1_waybel_0(X19,X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|~l1_orders_2(X18))&(r1_orders_2(X18,X21,esk7_4(X18,X19,X20,X21))|~r2_hidden(X21,X19)|~r2_hidden(X20,X19)|~m1_subset_1(X21,u1_struct_0(X18))|~m1_subset_1(X20,u1_struct_0(X18))|~v1_waybel_0(X19,X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|~l1_orders_2(X18)))))&((m1_subset_1(esk8_2(X18,X19),u1_struct_0(X18))|v1_waybel_0(X19,X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|~l1_orders_2(X18))&((m1_subset_1(esk9_2(X18,X19),u1_struct_0(X18))|v1_waybel_0(X19,X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|~l1_orders_2(X18))&((r2_hidden(esk8_2(X18,X19),X19)|v1_waybel_0(X19,X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|~l1_orders_2(X18))&((r2_hidden(esk9_2(X18,X19),X19)|v1_waybel_0(X19,X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|~l1_orders_2(X18))&(~m1_subset_1(X25,u1_struct_0(X18))|(~r2_hidden(X25,X19)|(~r1_orders_2(X18,esk8_2(X18,X19),X25)|~r1_orders_2(X18,esk9_2(X18,X19),X25)))|v1_waybel_0(X19,X18)|~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|~l1_orders_2(X18))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_waybel_0])])])])])).
cnf(c_0_8, plain, (r1_orders_2(X1,X2,X3)|X3!=k13_lattice3(X1,X2,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (m1_subset_1(k13_lattice3(X1,X2,X3),u1_struct_0(X1))|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (v1_waybel_0(X3,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r1_orders_2(X2,esk8_2(X2,X3),X1)|~r1_orders_2(X2,esk9_2(X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (r1_orders_2(X1,X2,k13_lattice3(X1,X2,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_8]), c_0_9])).
cnf(c_0_12, plain, (m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, plain, (r1_orders_2(X1,X2,X3)|X3!=k13_lattice3(X1,X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_14, plain, (v1_waybel_0(X1,X2)|~r1_orders_2(X2,esk8_2(X2,X1),k13_lattice3(X2,esk9_2(X2,X1),X3))|~r2_hidden(k13_lattice3(X2,esk9_2(X2,X1),X3),X1)|~m1_subset_1(k13_lattice3(X2,esk9_2(X2,X1),X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])).
cnf(c_0_15, plain, (r1_orders_2(X1,X2,k13_lattice3(X1,X3,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_13]), c_0_9])).
cnf(c_0_16, plain, (m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_17, negated_conjecture, ~(![X1]:((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))=>![X2]:((v12_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r2_hidden(X4,X2))=>r2_hidden(k13_lattice3(X1,X3,X4),X2)))))))), inference(assume_negation,[status(cth)],[t40_waybel_0])).
cnf(c_0_18, plain, (v1_waybel_0(X1,X2)|~r2_hidden(k13_lattice3(X2,esk9_2(X2,X1),esk8_2(X2,X1)),X1)|~m1_subset_1(k13_lattice3(X2,esk9_2(X2,X1),esk8_2(X2,X1)),u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_12]), c_0_16])).
fof(c_0_19, negated_conjecture, ![X10, X11]:((v5_orders_2(esk1_0)&(v1_lattice3(esk1_0)&l1_orders_2(esk1_0)))&((v12_waybel_0(esk2_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))&(((m1_subset_1(esk3_0,u1_struct_0(esk1_0))|~v1_waybel_0(esk2_0,esk1_0))&((m1_subset_1(esk4_0,u1_struct_0(esk1_0))|~v1_waybel_0(esk2_0,esk1_0))&(((r2_hidden(esk3_0,esk2_0)|~v1_waybel_0(esk2_0,esk1_0))&(r2_hidden(esk4_0,esk2_0)|~v1_waybel_0(esk2_0,esk1_0)))&(~r2_hidden(k13_lattice3(esk1_0,esk3_0,esk4_0),esk2_0)|~v1_waybel_0(esk2_0,esk1_0)))))&(v1_waybel_0(esk2_0,esk1_0)|(~m1_subset_1(X10,u1_struct_0(esk1_0))|(~m1_subset_1(X11,u1_struct_0(esk1_0))|(~r2_hidden(X10,esk2_0)|~r2_hidden(X11,esk2_0)|r2_hidden(k13_lattice3(esk1_0,X10,X11),esk2_0)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])])])).
cnf(c_0_20, plain, (v1_waybel_0(X1,X2)|~r2_hidden(k13_lattice3(X2,esk9_2(X2,X1),esk8_2(X2,X1)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_9]), c_0_12]), c_0_16])).
cnf(c_0_21, negated_conjecture, (v1_waybel_0(esk2_0,esk1_0)|r2_hidden(k13_lattice3(esk1_0,X1,X2),esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)|~r2_hidden(X2,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_22, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, negated_conjecture, (v1_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, negated_conjecture, (v1_waybel_0(esk2_0,esk1_0)|~r2_hidden(esk8_2(esk1_0,esk2_0),esk2_0)|~r2_hidden(esk9_2(esk1_0,esk2_0),esk2_0)|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25])])).
cnf(c_0_27, plain, (r2_hidden(esk9_2(X1,X2),X2)|v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_28, plain, ![X12, X13, X14, X15]:((~v12_waybel_0(X13,X12)|(~m1_subset_1(X14,u1_struct_0(X12))|(~m1_subset_1(X15,u1_struct_0(X12))|(~r2_hidden(X14,X13)|~r1_orders_2(X12,X15,X14)|r2_hidden(X15,X13))))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12))&((m1_subset_1(esk5_2(X12,X13),u1_struct_0(X12))|v12_waybel_0(X13,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12))&((m1_subset_1(esk6_2(X12,X13),u1_struct_0(X12))|v12_waybel_0(X13,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12))&(((r2_hidden(esk5_2(X12,X13),X13)|v12_waybel_0(X13,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12))&(r1_orders_2(X12,esk6_2(X12,X13),esk5_2(X12,X13))|v12_waybel_0(X13,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12)))&(~r2_hidden(esk6_2(X12,X13),X13)|v12_waybel_0(X13,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))|~l1_orders_2(X12)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d19_waybel_0])])])])])).
cnf(c_0_29, plain, (r1_orders_2(X2,X5,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,X3,X1)|~r1_orders_2(X2,X4,X1)|X5!=k13_lattice3(X2,X3,X4)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~v5_orders_2(X2)|~v1_lattice3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_30, negated_conjecture, (v1_waybel_0(esk2_0,esk1_0)|~r2_hidden(esk8_2(esk1_0,esk2_0),esk2_0)|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_22]), c_0_23])])).
cnf(c_0_31, plain, (r2_hidden(esk8_2(X1,X2),X2)|v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_32, plain, (r2_hidden(X4,X1)|~v12_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,X1)|~r1_orders_2(X2,X4,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33, plain, (r1_orders_2(X1,k13_lattice3(X1,X2,X3),X4)|~r1_orders_2(X1,X3,X4)|~r1_orders_2(X1,X2,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_29]), c_0_9])).
cnf(c_0_34, negated_conjecture, (v1_waybel_0(esk2_0,esk1_0)|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_22]), c_0_23])])).
cnf(c_0_35, plain, (r2_hidden(k13_lattice3(X1,X2,X3),X4)|~r1_orders_2(X1,X3,X5)|~r1_orders_2(X1,X2,X5)|~r2_hidden(X5,X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v12_waybel_0(X4,X1)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_9])).
cnf(c_0_36, plain, (r1_orders_2(X1,X2,esk7_4(X1,X3,X4,X2))|~r2_hidden(X2,X3)|~r2_hidden(X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v1_waybel_0(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_37, plain, (m1_subset_1(esk7_4(X1,X2,X3,X4),u1_struct_0(X1))|~r2_hidden(X4,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_38, negated_conjecture, (v1_waybel_0(esk2_0,esk1_0)|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_12]), c_0_22]), c_0_23])])).
cnf(c_0_39, plain, (r2_hidden(k13_lattice3(X1,X2,X3),X4)|~r1_orders_2(X1,X2,esk7_4(X1,X5,X6,X3))|~r2_hidden(esk7_4(X1,X5,X6,X3),X4)|~r2_hidden(X6,X5)|~r2_hidden(X3,X5)|~v1_waybel_0(X5,X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~v12_waybel_0(X4,X1)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37])).
cnf(c_0_40, plain, (r1_orders_2(X1,X2,esk7_4(X1,X3,X2,X4))|~r2_hidden(X4,X3)|~r2_hidden(X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_waybel_0(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_41, negated_conjecture, (~r2_hidden(k13_lattice3(esk1_0,esk3_0,esk4_0),esk2_0)|~v1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_42, negated_conjecture, (v1_waybel_0(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_16]), c_0_22]), c_0_23])])).
cnf(c_0_43, plain, (r2_hidden(k13_lattice3(X1,X2,X3),X4)|~r2_hidden(esk7_4(X1,X5,X2,X3),X4)|~r2_hidden(X2,X5)|~r2_hidden(X3,X5)|~v1_waybel_0(X5,X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v12_waybel_0(X4,X1)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_44, plain, (r2_hidden(esk7_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_45, negated_conjecture, (r2_hidden(esk3_0,esk2_0)|~v1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk4_0,esk2_0)|~v1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_47, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))|~v1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_48, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))|~v1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_49, negated_conjecture, (~r2_hidden(k13_lattice3(esk1_0,esk3_0,esk4_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_42])])).
cnf(c_0_50, plain, (r2_hidden(k13_lattice3(X1,X2,X3),X4)|~r2_hidden(X2,X4)|~r2_hidden(X3,X4)|~v1_waybel_0(X4,X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v12_waybel_0(X4,X1)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_51, negated_conjecture, (r2_hidden(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_42])])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk4_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_42])])).
cnf(c_0_53, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_42])])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_42])])).
cnf(c_0_55, negated_conjecture, (v12_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_56, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51]), c_0_52]), c_0_42]), c_0_22]), c_0_53]), c_0_54]), c_0_55]), c_0_23]), c_0_24]), c_0_25])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 57
# Proof object clause steps            : 46
# Proof object formula steps           : 11
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 51
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 34
# Processed clauses                    : 293
# ...of these trivial                  : 0
# ...subsumed                          : 132
# ...remaining for further processing  : 161
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 8
# Backward-rewritten                   : 8
# Generated clauses                    : 420
# ...of the previous two non-trivial   : 359
# Contextual simplify-reflections      : 62
# Paramodulations                      : 417
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 108
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 97
# Current number of unprocessed clauses: 126
# ...number of literals in the above   : 1598
# Current number of archived formulas  : 0
# Current number of archived clauses   : 50
# Clause-clause subsumption calls (NU) : 14626
# Rec. Clause-clause subsumption calls : 1295
# Non-unit clause-clause subsumptions  : 202
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 24771

# -------------------------------------------------
# User time                : 0.028 s
# System time              : 0.000 s
# Total time               : 0.028 s
# Maximum resident set size: 28720 pages
