# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', t2_xboole_1)).
fof(fraenkel_a_2_5_lattice3, axiom, ![X1, X2, X3]:((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))=>(r2_hidden(X1,a_2_5_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r3_lattices(X2,X4,X5)&r2_hidden(X5,X3))))))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', fraenkel_a_2_5_lattice3)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', t1_subset)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', d3_tarski)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', t1_xboole_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', t3_subset)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', d10_xboole_0)).
fof(t50_lattice3, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))&k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', t50_lattice3)).
fof(t3_xboole_1, axiom, ![X1]:(r1_tarski(X1,k1_xboole_0)=>X1=k1_xboole_0), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', t3_xboole_1)).
fof(d21_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X3=k15_lattice3(X1,X2)<=>(r4_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r4_lattice3(X1,X4,X2)=>r1_lattices(X1,X3,X4)))))))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', d21_lattice3)).
fof(dt_k15_lattice3, axiom, ![X1, X2]:((~(v2_struct_0(X1))&l3_lattices(X1))=>m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', dt_k15_lattice3)).
fof(t47_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(k15_lattice3(X1,X2)=k15_lattice3(X1,a_2_5_lattice3(X1,X2))&k16_lattice3(X1,X2)=k16_lattice3(X1,a_2_6_lattice3(X1,X2)))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', t47_lattice3)).
fof(t26_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&(v4_lattices(X1)&l2_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattices(X1,X2,X3)&r1_lattices(X1,X3,X2))=>X2=X3)))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', t26_lattices)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', cc1_lattices)).
fof(dt_l3_lattices, axiom, ![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', dt_l3_lattices)).
fof(t45_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(X2=k15_lattice3(X1,a_2_3_lattice3(X1,X2))&X2=k16_lattice3(X1,a_2_4_lattice3(X1,X2))))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', t45_lattice3)).
fof(d17_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r4_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X4,X2)))))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', d17_lattice3)).
fof(fraenkel_a_2_3_lattice3, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))&m1_subset_1(X3,u1_struct_0(X2)))=>(r2_hidden(X1,a_2_3_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&r3_lattices(X2,X4,X3))))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', fraenkel_a_2_3_lattice3)).
fof(t46_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(r1_tarski(X2,X3)=>(r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3))&r3_lattices(X1,k16_lattice3(X1,X3),k16_lattice3(X1,X2))))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', t46_lattice3)).
fof(d9_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>(v9_lattices(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_lattices(X1,X2,k1_lattices(X1,X2,X3))=X2)))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', d9_lattices)).
fof(d3_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l2_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattices(X1,X2,X3)<=>k1_lattices(X1,X2,X3)=X3)))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', d3_lattices)).
fof(d16_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v13_lattices(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(X2=k5_lattices(X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_lattices(X1,X2,X3)=X2&k2_lattices(X1,X3,X2)=X2)))))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', d16_lattices)).
fof(dt_k5_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>m1_subset_1(k5_lattices(X1),u1_struct_0(X1))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', dt_k5_lattices)).
fof(d6_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v6_lattices(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2))))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', d6_lattices)).
fof(d13_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v13_lattices(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_lattices(X1,X2,X3)=X2&k2_lattices(X1,X3,X2)=X2))))), file('data/MPTP/proofs/7cbffc0b72df4ca7ae712b1b93fbeb5d.E_input', d13_lattices)).
fof(c_0_26, plain, ![X6, X7]:(~r2_hidden(X6,X7)|~r1_tarski(X7,X6)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_27, plain, ![X8]:r1_tarski(k1_xboole_0,X8), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_28, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_29, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_30, plain, ![X9, X10, X11, X14, X15, X16]:(((m1_subset_1(esk1_3(X9,X10,X11),u1_struct_0(X10))|~r2_hidden(X9,a_2_5_lattice3(X10,X11))|(v2_struct_0(X10)|(~v10_lattices(X10)|(~v4_lattice3(X10)|~l3_lattices(X10)))))&((X9=esk1_3(X9,X10,X11)|~r2_hidden(X9,a_2_5_lattice3(X10,X11))|(v2_struct_0(X10)|(~v10_lattices(X10)|(~v4_lattice3(X10)|~l3_lattices(X10)))))&((m1_subset_1(esk2_3(X9,X10,X11),u1_struct_0(X10))|~r2_hidden(X9,a_2_5_lattice3(X10,X11))|(v2_struct_0(X10)|(~v10_lattices(X10)|(~v4_lattice3(X10)|~l3_lattices(X10)))))&((r3_lattices(X10,esk1_3(X9,X10,X11),esk2_3(X9,X10,X11))|~r2_hidden(X9,a_2_5_lattice3(X10,X11))|(v2_struct_0(X10)|(~v10_lattices(X10)|(~v4_lattice3(X10)|~l3_lattices(X10)))))&(r2_hidden(esk2_3(X9,X10,X11),X11)|~r2_hidden(X9,a_2_5_lattice3(X10,X11))|(v2_struct_0(X10)|(~v10_lattices(X10)|(~v4_lattice3(X10)|~l3_lattices(X10)))))))))&(~m1_subset_1(X15,u1_struct_0(X10))|(X9!=X15|(~m1_subset_1(X16,u1_struct_0(X10))|(~r3_lattices(X10,X15,X16)|~r2_hidden(X16,X14))))|r2_hidden(X9,a_2_5_lattice3(X10,X14))|(v2_struct_0(X10)|(~v10_lattices(X10)|(~v4_lattice3(X10)|~l3_lattices(X10)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_5_lattice3])])])])])])])).
fof(c_0_31, plain, ![X17, X18]:(~r2_hidden(X17,X18)|m1_subset_1(X17,X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
fof(c_0_32, plain, ![X19, X20, X21, X22, X23]:((~r1_tarski(X19,X20)|(~r2_hidden(X21,X19)|r2_hidden(X21,X20)))&((r2_hidden(esk3_2(X22,X23),X22)|r1_tarski(X22,X23))&(~r2_hidden(esk3_2(X22,X23),X23)|r1_tarski(X22,X23)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_33, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34, plain, (r2_hidden(esk2_3(X1,X2,X3),X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_5_lattice3(X2,X3))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_35, plain, ![X25, X26, X27]:(~r1_tarski(X25,X26)|~r1_tarski(X26,X27)|r1_tarski(X25,X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_36, plain, ![X28, X29]:((~m1_subset_1(X28,k1_zfmisc_1(X29))|r1_tarski(X28,X29))&(~r1_tarski(X28,X29)|m1_subset_1(X28,k1_zfmisc_1(X29)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_37, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_39, plain, ![X30, X31]:(((r1_tarski(X30,X31)|X30!=X31)&(r1_tarski(X31,X30)|X30!=X31))&(~r1_tarski(X30,X31)|~r1_tarski(X31,X30)|X30=X31)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_40, plain, (v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(X2,a_2_5_lattice3(X1,k1_xboole_0))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_41, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_42, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_43, plain, (m1_subset_1(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_44, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_45, plain, (v2_struct_0(X1)|r1_tarski(a_2_5_lattice3(X1,k1_xboole_0),X2)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_40, c_0_38])).
cnf(c_0_46, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_41, c_0_29])).
cnf(c_0_47, plain, (r1_tarski(esk3_2(k1_zfmisc_1(X1),X2),X1)|r1_tarski(k1_zfmisc_1(X1),X2)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_48, plain, (X1=a_2_5_lattice3(X2,k1_xboole_0)|v2_struct_0(X2)|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)|~r1_tarski(X1,a_2_5_lattice3(X2,k1_xboole_0))), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_49, plain, (r1_tarski(esk3_2(k1_zfmisc_1(k1_xboole_0),X1),X2)|r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
fof(c_0_50, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))&k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0)))), inference(assume_negation,[status(cth)],[t50_lattice3])).
fof(c_0_51, plain, ![X33]:(~r1_tarski(X33,k1_xboole_0)|X33=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_xboole_1])])).
cnf(c_0_52, plain, (esk3_2(k1_zfmisc_1(k1_xboole_0),X1)=a_2_5_lattice3(X2,k1_xboole_0)|v2_struct_0(X2)|r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
fof(c_0_53, negated_conjecture, ((~v2_struct_0(esk4_0)&(v10_lattices(esk4_0)&(v4_lattice3(esk4_0)&l3_lattices(esk4_0))))&(v2_struct_0(esk4_0)|(~v10_lattices(esk4_0)|(~v13_lattices(esk4_0)|~l3_lattices(esk4_0)))|k5_lattices(esk4_0)!=k15_lattice3(esk4_0,k1_xboole_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_50])])])])).
cnf(c_0_54, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_55, plain, (v2_struct_0(X1)|r1_tarski(k1_zfmisc_1(k1_xboole_0),X2)|r2_hidden(a_2_5_lattice3(X1,k1_xboole_0),k1_zfmisc_1(k1_xboole_0))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_38, c_0_52])).
cnf(c_0_56, negated_conjecture, (l3_lattices(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_57, negated_conjecture, (v4_lattice3(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_58, negated_conjecture, (v10_lattices(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_59, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_60, plain, (esk3_2(k1_zfmisc_1(k1_xboole_0),X1)=k1_xboole_0|r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_54, c_0_47])).
cnf(c_0_61, negated_conjecture, (r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)|r2_hidden(a_2_5_lattice3(esk4_0,k1_xboole_0),k1_zfmisc_1(k1_xboole_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_62, plain, (r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_38, c_0_60])).
cnf(c_0_63, negated_conjecture, (r2_hidden(a_2_5_lattice3(esk4_0,k1_xboole_0),k1_zfmisc_1(k1_xboole_0))|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_28, c_0_61])).
cnf(c_0_64, plain, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_54, c_0_62])).
fof(c_0_65, plain, ![X34, X35, X36, X37]:(((r4_lattice3(X34,X36,X35)|X36!=k15_lattice3(X34,X35)|~m1_subset_1(X36,u1_struct_0(X34))|(v2_struct_0(X34)|(~v10_lattices(X34)|(~v4_lattice3(X34)|~l3_lattices(X34))))|(v2_struct_0(X34)|~l3_lattices(X34)))&(~m1_subset_1(X37,u1_struct_0(X34))|(~r4_lattice3(X34,X37,X35)|r1_lattices(X34,X36,X37))|X36!=k15_lattice3(X34,X35)|~m1_subset_1(X36,u1_struct_0(X34))|(v2_struct_0(X34)|(~v10_lattices(X34)|(~v4_lattice3(X34)|~l3_lattices(X34))))|(v2_struct_0(X34)|~l3_lattices(X34))))&((m1_subset_1(esk5_3(X34,X35,X36),u1_struct_0(X34))|~r4_lattice3(X34,X36,X35)|X36=k15_lattice3(X34,X35)|~m1_subset_1(X36,u1_struct_0(X34))|(v2_struct_0(X34)|(~v10_lattices(X34)|(~v4_lattice3(X34)|~l3_lattices(X34))))|(v2_struct_0(X34)|~l3_lattices(X34)))&((r4_lattice3(X34,esk5_3(X34,X35,X36),X35)|~r4_lattice3(X34,X36,X35)|X36=k15_lattice3(X34,X35)|~m1_subset_1(X36,u1_struct_0(X34))|(v2_struct_0(X34)|(~v10_lattices(X34)|(~v4_lattice3(X34)|~l3_lattices(X34))))|(v2_struct_0(X34)|~l3_lattices(X34)))&(~r1_lattices(X34,X36,esk5_3(X34,X35,X36))|~r4_lattice3(X34,X36,X35)|X36=k15_lattice3(X34,X35)|~m1_subset_1(X36,u1_struct_0(X34))|(v2_struct_0(X34)|(~v10_lattices(X34)|(~v4_lattice3(X34)|~l3_lattices(X34))))|(v2_struct_0(X34)|~l3_lattices(X34)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d21_lattice3])])])])])])).
cnf(c_0_66, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|r2_hidden(a_2_5_lattice3(esk4_0,k1_xboole_0),k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_67, plain, (r1_lattices(X2,X4,X1)|v2_struct_0(X2)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r4_lattice3(X2,X1,X3)|X4!=k15_lattice3(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
fof(c_0_68, plain, ![X39, X40]:(v2_struct_0(X39)|~l3_lattices(X39)|m1_subset_1(k15_lattice3(X39,X40),u1_struct_0(X39))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k15_lattice3])])])).
cnf(c_0_69, plain, (X1=k1_zfmisc_1(k1_xboole_0)|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))|~r1_tarski(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_44, c_0_62])).
cnf(c_0_70, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|m1_subset_1(a_2_5_lattice3(esk4_0,k1_xboole_0),k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_37, c_0_66])).
cnf(c_0_71, plain, (v2_struct_0(X2)|r1_lattices(X2,X4,X1)|X4!=k15_lattice3(X2,X3)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)|~r4_lattice3(X2,X1,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))), inference(cn,[status(thm)],[c_0_67])).
cnf(c_0_72, plain, (v2_struct_0(X1)|m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
fof(c_0_73, plain, ![X41, X42]:((k15_lattice3(X41,X42)=k15_lattice3(X41,a_2_5_lattice3(X41,X42))|(v2_struct_0(X41)|(~v10_lattices(X41)|(~v4_lattice3(X41)|~l3_lattices(X41)))))&(k16_lattice3(X41,X42)=k16_lattice3(X41,a_2_6_lattice3(X41,X42))|(v2_struct_0(X41)|(~v10_lattices(X41)|(~v4_lattice3(X41)|~l3_lattices(X41)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t47_lattice3])])])])])).
cnf(c_0_74, plain, (a_2_5_lattice3(X1,k1_xboole_0)=k1_zfmisc_1(k1_xboole_0)|v2_struct_0(X1)|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_69, c_0_45])).
cnf(c_0_75, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|r1_tarski(a_2_5_lattice3(esk4_0,k1_xboole_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_42, c_0_70])).
fof(c_0_76, plain, ![X43, X44, X45]:(v2_struct_0(X43)|(~v4_lattices(X43)|~l2_lattices(X43))|(~m1_subset_1(X44,u1_struct_0(X43))|(~m1_subset_1(X45,u1_struct_0(X43))|(~r1_lattices(X43,X44,X45)|~r1_lattices(X43,X45,X44)|X44=X45)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t26_lattices])])])])).
fof(c_0_77, plain, ![X46]:((~v2_struct_0(X46)|(v2_struct_0(X46)|~v10_lattices(X46))|~l3_lattices(X46))&((v4_lattices(X46)|(v2_struct_0(X46)|~v10_lattices(X46))|~l3_lattices(X46))&((v5_lattices(X46)|(v2_struct_0(X46)|~v10_lattices(X46))|~l3_lattices(X46))&((v6_lattices(X46)|(v2_struct_0(X46)|~v10_lattices(X46))|~l3_lattices(X46))&((v7_lattices(X46)|(v2_struct_0(X46)|~v10_lattices(X46))|~l3_lattices(X46))&((v8_lattices(X46)|(v2_struct_0(X46)|~v10_lattices(X46))|~l3_lattices(X46))&(v9_lattices(X46)|(v2_struct_0(X46)|~v10_lattices(X46))|~l3_lattices(X46)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
fof(c_0_78, plain, ![X47]:((l1_lattices(X47)|~l3_lattices(X47))&(l2_lattices(X47)|~l3_lattices(X47))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
cnf(c_0_79, plain, (r1_lattices(X1,k15_lattice3(X1,X2),X3)|v2_struct_0(X1)|~r4_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_71]), c_0_72])).
cnf(c_0_80, plain, (k15_lattice3(X1,X2)=k15_lattice3(X1,a_2_5_lattice3(X1,X2))|v2_struct_0(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_81, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)=a_2_5_lattice3(esk4_0,k1_xboole_0)|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_82, negated_conjecture, (a_2_5_lattice3(esk4_0,k1_xboole_0)=k1_xboole_0|k1_zfmisc_1(k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_75]), c_0_29])])).
cnf(c_0_83, plain, (r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X1)|X2!=k15_lattice3(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_84, plain, (v2_struct_0(X1)|X2=X3|~v4_lattices(X1)|~l2_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattices(X1,X2,X3)|~r1_lattices(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_85, plain, (v4_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_86, plain, (l2_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_87, plain, (r1_lattices(X1,k15_lattice3(X1,X2),X3)|v2_struct_0(X1)|~r4_lattice3(X1,X3,a_2_5_lattice3(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_88, negated_conjecture, (a_2_5_lattice3(esk4_0,k1_xboole_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_33])).
cnf(c_0_89, plain, (v2_struct_0(X1)|r4_lattice3(X1,X2,X3)|X2!=k15_lattice3(X1,X3)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_83])).
fof(c_0_90, plain, ![X48, X49]:((X49=k15_lattice3(X48,a_2_3_lattice3(X48,X49))|~m1_subset_1(X49,u1_struct_0(X48))|(v2_struct_0(X48)|(~v10_lattices(X48)|(~v4_lattice3(X48)|~l3_lattices(X48)))))&(X49=k16_lattice3(X48,a_2_4_lattice3(X48,X49))|~m1_subset_1(X49,u1_struct_0(X48))|(v2_struct_0(X48)|(~v10_lattices(X48)|(~v4_lattice3(X48)|~l3_lattices(X48)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t45_lattice3])])])])])).
cnf(c_0_91, plain, (X1=k15_lattice3(X2,X3)|v2_struct_0(X2)|~r1_lattices(X2,X1,k15_lattice3(X2,X3))|~r4_lattice3(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_79]), c_0_72]), c_0_85]), c_0_86])).
cnf(c_0_92, negated_conjecture, (r1_lattices(esk4_0,k15_lattice3(esk4_0,k1_xboole_0),X1)|~r4_lattice3(esk4_0,X1,k1_xboole_0)|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_93, plain, (r4_lattice3(X1,k15_lattice3(X1,X2),X2)|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_89]), c_0_72])).
cnf(c_0_94, plain, (X1=k15_lattice3(X2,a_2_3_lattice3(X2,X1))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_90])).
fof(c_0_95, plain, ![X50, X51, X52, X53, X54]:((~r4_lattice3(X50,X51,X52)|(~m1_subset_1(X53,u1_struct_0(X50))|(~r2_hidden(X53,X52)|r1_lattices(X50,X53,X51)))|~m1_subset_1(X51,u1_struct_0(X50))|(v2_struct_0(X50)|~l3_lattices(X50)))&((m1_subset_1(esk6_3(X50,X51,X54),u1_struct_0(X50))|r4_lattice3(X50,X51,X54)|~m1_subset_1(X51,u1_struct_0(X50))|(v2_struct_0(X50)|~l3_lattices(X50)))&((r2_hidden(esk6_3(X50,X51,X54),X54)|r4_lattice3(X50,X51,X54)|~m1_subset_1(X51,u1_struct_0(X50))|(v2_struct_0(X50)|~l3_lattices(X50)))&(~r1_lattices(X50,esk6_3(X50,X51,X54),X51)|r4_lattice3(X50,X51,X54)|~m1_subset_1(X51,u1_struct_0(X50))|(v2_struct_0(X50)|~l3_lattices(X50)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_lattice3])])])])])])])).
cnf(c_0_96, negated_conjecture, (k15_lattice3(esk4_0,k1_xboole_0)=k15_lattice3(esk4_0,X1)|~r4_lattice3(esk4_0,k15_lattice3(esk4_0,k1_xboole_0),X1)|~r4_lattice3(esk4_0,k15_lattice3(esk4_0,X1),k1_xboole_0)|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))|~m1_subset_1(k15_lattice3(esk4_0,X1),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_97, plain, (r4_lattice3(X1,X2,a_2_3_lattice3(X1,X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_98, plain, (r2_hidden(esk6_3(X1,X2,X3),X3)|r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_99, negated_conjecture, (k15_lattice3(esk4_0,a_2_3_lattice3(esk4_0,k15_lattice3(esk4_0,k1_xboole_0)))=k15_lattice3(esk4_0,k1_xboole_0)|~r4_lattice3(esk4_0,k15_lattice3(esk4_0,a_2_3_lattice3(esk4_0,k15_lattice3(esk4_0,k1_xboole_0))),k1_xboole_0)|~m1_subset_1(k15_lattice3(esk4_0,a_2_3_lattice3(esk4_0,k15_lattice3(esk4_0,k1_xboole_0))),u1_struct_0(esk4_0))|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_100, plain, (r4_lattice3(X1,X2,k1_xboole_0)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_33, c_0_98])).
cnf(c_0_101, plain, (r1_lattices(X1,X4,X2)|v2_struct_0(X1)|~r4_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_102, negated_conjecture, (k15_lattice3(esk4_0,a_2_3_lattice3(esk4_0,k15_lattice3(esk4_0,k1_xboole_0)))=k15_lattice3(esk4_0,k1_xboole_0)|~m1_subset_1(k15_lattice3(esk4_0,a_2_3_lattice3(esk4_0,k15_lattice3(esk4_0,k1_xboole_0))),u1_struct_0(esk4_0))|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_56])]), c_0_59])).
fof(c_0_103, plain, ![X56, X57, X58, X60]:(((m1_subset_1(esk7_3(X56,X57,X58),u1_struct_0(X57))|~r2_hidden(X56,a_2_3_lattice3(X57,X58))|(v2_struct_0(X57)|(~v10_lattices(X57)|(~v4_lattice3(X57)|~l3_lattices(X57)))|~m1_subset_1(X58,u1_struct_0(X57))))&((X56=esk7_3(X56,X57,X58)|~r2_hidden(X56,a_2_3_lattice3(X57,X58))|(v2_struct_0(X57)|(~v10_lattices(X57)|(~v4_lattice3(X57)|~l3_lattices(X57)))|~m1_subset_1(X58,u1_struct_0(X57))))&(r3_lattices(X57,esk7_3(X56,X57,X58),X58)|~r2_hidden(X56,a_2_3_lattice3(X57,X58))|(v2_struct_0(X57)|(~v10_lattices(X57)|(~v4_lattice3(X57)|~l3_lattices(X57)))|~m1_subset_1(X58,u1_struct_0(X57))))))&(~m1_subset_1(X60,u1_struct_0(X57))|(X56!=X60|~r3_lattices(X57,X60,X58))|r2_hidden(X56,a_2_3_lattice3(X57,X58))|(v2_struct_0(X57)|(~v10_lattices(X57)|(~v4_lattice3(X57)|~l3_lattices(X57)))|~m1_subset_1(X58,u1_struct_0(X57))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_3_lattice3])])])])])])).
cnf(c_0_104, plain, (r1_lattices(X1,X2,k15_lattice3(X1,X3))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~r2_hidden(X2,X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_93]), c_0_72])).
cnf(c_0_105, negated_conjecture, (k15_lattice3(esk4_0,a_2_3_lattice3(esk4_0,k15_lattice3(esk4_0,k1_xboole_0)))=k15_lattice3(esk4_0,k1_xboole_0)|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_72]), c_0_56])]), c_0_59])).
cnf(c_0_106, plain, (r2_hidden(X3,a_2_3_lattice3(X2,X4))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~r3_lattices(X2,X1,X4)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_107, negated_conjecture, (r1_lattices(esk4_0,X1,k15_lattice3(esk4_0,k1_xboole_0))|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))|~m1_subset_1(X1,u1_struct_0(esk4_0))|~r2_hidden(X1,a_2_3_lattice3(esk4_0,k15_lattice3(esk4_0,k1_xboole_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_108, plain, (v2_struct_0(X1)|r2_hidden(X2,a_2_3_lattice3(X1,X3))|~r3_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(er,[status(thm)],[c_0_106])).
fof(c_0_109, plain, ![X61, X62, X63]:((r3_lattices(X61,k15_lattice3(X61,X62),k15_lattice3(X61,X63))|~r1_tarski(X62,X63)|(v2_struct_0(X61)|(~v10_lattices(X61)|(~v4_lattice3(X61)|~l3_lattices(X61)))))&(r3_lattices(X61,k16_lattice3(X61,X63),k16_lattice3(X61,X62))|~r1_tarski(X62,X63)|(v2_struct_0(X61)|(~v10_lattices(X61)|(~v4_lattice3(X61)|~l3_lattices(X61)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t46_lattice3])])])])])).
cnf(c_0_110, negated_conjecture, (r1_lattices(esk4_0,X1,k15_lattice3(esk4_0,k1_xboole_0))|~r3_lattices(esk4_0,X1,k15_lattice3(esk4_0,k1_xboole_0))|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_111, plain, (r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3))|v2_struct_0(X1)|~r1_tarski(X2,X3)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_109])).
cnf(c_0_112, negated_conjecture, (r1_lattices(esk4_0,k15_lattice3(esk4_0,X1),k15_lattice3(esk4_0,k1_xboole_0))|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))|~m1_subset_1(k15_lattice3(esk4_0,X1),u1_struct_0(esk4_0))|~r1_tarski(X1,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_113, negated_conjecture, (k15_lattice3(esk4_0,X1)=k15_lattice3(esk4_0,k1_xboole_0)|~r4_lattice3(esk4_0,k15_lattice3(esk4_0,X1),k1_xboole_0)|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))|~m1_subset_1(k15_lattice3(esk4_0,X1),u1_struct_0(esk4_0))|~r1_tarski(X1,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_112]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_114, negated_conjecture, (k15_lattice3(esk4_0,X1)=k15_lattice3(esk4_0,k1_xboole_0)|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))|~m1_subset_1(k15_lattice3(esk4_0,X1),u1_struct_0(esk4_0))|~r1_tarski(X1,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_100]), c_0_56])]), c_0_59])).
cnf(c_0_115, negated_conjecture, (k15_lattice3(esk4_0,X1)=k15_lattice3(esk4_0,k1_xboole_0)|~m1_subset_1(k15_lattice3(esk4_0,X1),u1_struct_0(esk4_0))|~r1_tarski(X1,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_72]), c_0_56])]), c_0_59])).
fof(c_0_116, plain, ![X64, X65, X66]:((~v9_lattices(X64)|(~m1_subset_1(X65,u1_struct_0(X64))|(~m1_subset_1(X66,u1_struct_0(X64))|k2_lattices(X64,X65,k1_lattices(X64,X65,X66))=X65))|(v2_struct_0(X64)|~l3_lattices(X64)))&((m1_subset_1(esk8_1(X64),u1_struct_0(X64))|v9_lattices(X64)|(v2_struct_0(X64)|~l3_lattices(X64)))&((m1_subset_1(esk9_1(X64),u1_struct_0(X64))|v9_lattices(X64)|(v2_struct_0(X64)|~l3_lattices(X64)))&(k2_lattices(X64,esk8_1(X64),k1_lattices(X64,esk8_1(X64),esk9_1(X64)))!=esk8_1(X64)|v9_lattices(X64)|(v2_struct_0(X64)|~l3_lattices(X64)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d9_lattices])])])])])])).
fof(c_0_117, plain, ![X69, X70, X71]:((~r1_lattices(X69,X70,X71)|k1_lattices(X69,X70,X71)=X71|~m1_subset_1(X71,u1_struct_0(X69))|~m1_subset_1(X70,u1_struct_0(X69))|(v2_struct_0(X69)|~l2_lattices(X69)))&(k1_lattices(X69,X70,X71)!=X71|r1_lattices(X69,X70,X71)|~m1_subset_1(X71,u1_struct_0(X69))|~m1_subset_1(X70,u1_struct_0(X69))|(v2_struct_0(X69)|~l2_lattices(X69)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d3_lattices])])])])])).
cnf(c_0_118, plain, (r4_lattice3(X1,X2,a_2_5_lattice3(X3,k1_xboole_0))|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X3)|~l3_lattices(X1)|~v4_lattice3(X3)|~v10_lattices(X3)), inference(spm,[status(thm)],[c_0_40, c_0_98])).
cnf(c_0_119, negated_conjecture, (k15_lattice3(esk4_0,X1)=k15_lattice3(esk4_0,k1_xboole_0)|~r1_tarski(X1,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_72]), c_0_56])]), c_0_59])).
fof(c_0_120, plain, ![X72, X73, X74]:(((k2_lattices(X72,X73,X74)=X73|~m1_subset_1(X74,u1_struct_0(X72))|X73!=k5_lattices(X72)|~m1_subset_1(X73,u1_struct_0(X72))|~v13_lattices(X72)|(v2_struct_0(X72)|~l1_lattices(X72)))&(k2_lattices(X72,X74,X73)=X73|~m1_subset_1(X74,u1_struct_0(X72))|X73!=k5_lattices(X72)|~m1_subset_1(X73,u1_struct_0(X72))|~v13_lattices(X72)|(v2_struct_0(X72)|~l1_lattices(X72))))&((m1_subset_1(esk10_2(X72,X73),u1_struct_0(X72))|X73=k5_lattices(X72)|~m1_subset_1(X73,u1_struct_0(X72))|~v13_lattices(X72)|(v2_struct_0(X72)|~l1_lattices(X72)))&(k2_lattices(X72,X73,esk10_2(X72,X73))!=X73|k2_lattices(X72,esk10_2(X72,X73),X73)!=X73|X73=k5_lattices(X72)|~m1_subset_1(X73,u1_struct_0(X72))|~v13_lattices(X72)|(v2_struct_0(X72)|~l1_lattices(X72))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattices])])])])])])).
fof(c_0_121, plain, ![X76]:(v2_struct_0(X76)|~l1_lattices(X76)|m1_subset_1(k5_lattices(X76),u1_struct_0(X76))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_lattices])])])).
cnf(c_0_122, plain, (k2_lattices(X1,X2,k1_lattices(X1,X2,X3))=X2|v2_struct_0(X1)|~v9_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_116])).
cnf(c_0_123, plain, (k1_lattices(X1,X2,X3)=X3|v2_struct_0(X1)|~r1_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_117])).
cnf(c_0_124, plain, (r1_lattices(X1,k15_lattice3(X1,k1_xboole_0),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_87, c_0_118])).
cnf(c_0_125, negated_conjecture, (m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))|~r1_tarski(X1,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_119]), c_0_56])]), c_0_59])).
cnf(c_0_126, plain, (k2_lattices(X1,X2,X3)=X3|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|X3!=k5_lattices(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v13_lattices(X1)|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_127, plain, (v2_struct_0(X1)|m1_subset_1(k5_lattices(X1),u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_128, plain, (k2_lattices(X1,X2,X3)=X2|v2_struct_0(X1)|~v9_lattices(X1)|~r1_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_86])).
cnf(c_0_129, negated_conjecture, (r1_lattices(esk4_0,k15_lattice3(esk4_0,k1_xboole_0),X1)|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_119]), c_0_56]), c_0_57]), c_0_58]), c_0_29])]), c_0_59])).
cnf(c_0_130, negated_conjecture, (m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))), inference(spm,[status(thm)],[c_0_125, c_0_29])).
cnf(c_0_131, negated_conjecture, (v2_struct_0(esk4_0)|~v10_lattices(esk4_0)|~v13_lattices(esk4_0)|~l3_lattices(esk4_0)|k5_lattices(esk4_0)!=k15_lattice3(esk4_0,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_53])).
fof(c_0_132, plain, ![X77, X78, X79]:((~v6_lattices(X77)|(~m1_subset_1(X78,u1_struct_0(X77))|(~m1_subset_1(X79,u1_struct_0(X77))|k2_lattices(X77,X78,X79)=k2_lattices(X77,X79,X78)))|(v2_struct_0(X77)|~l1_lattices(X77)))&((m1_subset_1(esk11_1(X77),u1_struct_0(X77))|v6_lattices(X77)|(v2_struct_0(X77)|~l1_lattices(X77)))&((m1_subset_1(esk12_1(X77),u1_struct_0(X77))|v6_lattices(X77)|(v2_struct_0(X77)|~l1_lattices(X77)))&(k2_lattices(X77,esk11_1(X77),esk12_1(X77))!=k2_lattices(X77,esk12_1(X77),esk11_1(X77))|v6_lattices(X77)|(v2_struct_0(X77)|~l1_lattices(X77)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_lattices])])])])])])).
cnf(c_0_133, plain, (k2_lattices(X1,X2,k5_lattices(X1))=k5_lattices(X1)|v2_struct_0(X1)|~l1_lattices(X1)|~v13_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_126]), c_0_127])).
cnf(c_0_134, negated_conjecture, (k2_lattices(esk4_0,k15_lattice3(esk4_0,k1_xboole_0),X1)=k15_lattice3(esk4_0,k1_xboole_0)|~v9_lattices(esk4_0)|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129]), c_0_130]), c_0_56])]), c_0_59])).
cnf(c_0_135, negated_conjecture, (k5_lattices(esk4_0)!=k15_lattice3(esk4_0,k1_xboole_0)|~v13_lattices(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_131, c_0_58]), c_0_56])]), c_0_59])).
cnf(c_0_136, plain, (k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2)|v2_struct_0(X1)|~v6_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_137, plain, (v6_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_138, plain, (l1_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_139, negated_conjecture, (~l1_lattices(esk4_0)|~v9_lattices(esk4_0)|~v13_lattices(esk4_0)|~m1_subset_1(k5_lattices(esk4_0),u1_struct_0(esk4_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134]), c_0_130])]), c_0_59]), c_0_135])).
fof(c_0_140, plain, ![X82, X84, X85]:(((m1_subset_1(esk13_1(X82),u1_struct_0(X82))|~v13_lattices(X82)|(v2_struct_0(X82)|~l1_lattices(X82)))&((k2_lattices(X82,esk13_1(X82),X84)=esk13_1(X82)|~m1_subset_1(X84,u1_struct_0(X82))|~v13_lattices(X82)|(v2_struct_0(X82)|~l1_lattices(X82)))&(k2_lattices(X82,X84,esk13_1(X82))=esk13_1(X82)|~m1_subset_1(X84,u1_struct_0(X82))|~v13_lattices(X82)|(v2_struct_0(X82)|~l1_lattices(X82)))))&((m1_subset_1(esk14_2(X82,X85),u1_struct_0(X82))|~m1_subset_1(X85,u1_struct_0(X82))|v13_lattices(X82)|(v2_struct_0(X82)|~l1_lattices(X82)))&(k2_lattices(X82,X85,esk14_2(X82,X85))!=X85|k2_lattices(X82,esk14_2(X82,X85),X85)!=X85|~m1_subset_1(X85,u1_struct_0(X82))|v13_lattices(X82)|(v2_struct_0(X82)|~l1_lattices(X82))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_lattices])])])])])])).
cnf(c_0_141, plain, (k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_137]), c_0_138])).
cnf(c_0_142, negated_conjecture, (~l1_lattices(esk4_0)|~v9_lattices(esk4_0)|~v13_lattices(esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_127]), c_0_59])).
cnf(c_0_143, plain, (v9_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_144, plain, (v13_lattices(X1)|v2_struct_0(X1)|k2_lattices(X1,X2,esk14_2(X1,X2))!=X2|k2_lattices(X1,esk14_2(X1,X2),X2)!=X2|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_140])).
cnf(c_0_145, negated_conjecture, (k2_lattices(esk4_0,X1,k15_lattice3(esk4_0,k1_xboole_0))=k15_lattice3(esk4_0,k1_xboole_0)|~v9_lattices(esk4_0)|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_134]), c_0_130]), c_0_56]), c_0_58])]), c_0_59])).
cnf(c_0_146, negated_conjecture, (~l1_lattices(esk4_0)|~v13_lattices(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_143]), c_0_56]), c_0_58])]), c_0_59])).
cnf(c_0_147, negated_conjecture, (~l1_lattices(esk4_0)|~v9_lattices(esk4_0)|~m1_subset_1(esk14_2(esk4_0,k15_lattice3(esk4_0,k1_xboole_0)),u1_struct_0(esk4_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_145]), c_0_130])]), c_0_59]), c_0_134]), c_0_146])).
cnf(c_0_148, plain, (m1_subset_1(esk14_2(X1,X2),u1_struct_0(X1))|v13_lattices(X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_140])).
cnf(c_0_149, negated_conjecture, (~l1_lattices(esk4_0)|~v9_lattices(esk4_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_148]), c_0_130])]), c_0_59]), c_0_146])).
cnf(c_0_150, negated_conjecture, (~l1_lattices(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_143]), c_0_56]), c_0_58])]), c_0_59])).
cnf(c_0_151, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_138]), c_0_56])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 152
# Proof object clause steps            : 99
# Proof object formula steps           : 53
# Proof object conjectures             : 41
# Proof object clause conjectures      : 38
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 36
# Proof object initial formulas used   : 26
# Proof object generating inferences   : 56
# Proof object simplifying inferences  : 122
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 26
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 74
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 73
# Processed clauses                    : 2087
# ...of these trivial                  : 15
# ...subsumed                          : 1431
# ...remaining for further processing  : 641
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 30
# Backward-rewritten                   : 56
# Generated clauses                    : 6607
# ...of the previous two non-trivial   : 6090
# Contextual simplify-reflections      : 34
# Paramodulations                      : 6599
# Factorizations                       : 0
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 475
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 457
# Current number of unprocessed clauses: 4004
# ...number of literals in the above   : 27051
# Current number of archived formulas  : 0
# Current number of archived clauses   : 158
# Clause-clause subsumption calls (NU) : 81107
# Rec. Clause-clause subsumption calls : 13660
# Non-unit clause-clause subsumptions  : 1339
# Unit Clause-clause subsumption calls : 99
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 15
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 160378

# -------------------------------------------------
# User time                : 0.116 s
# System time              : 0.003 s
# Total time               : 0.120 s
# Maximum resident set size: 31544 pages
