# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/7d1ac91fca294d128bb8aed501eb880b.E_input', d3_tarski)).
fof(t45_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(r2_hidden(X3,u1_struct_0(X1))=>(r2_hidden(X3,k2_pre_topc(X1,X2))<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>((v4_pre_topc(X4,X1)&r1_tarski(X2,X4))=>r2_hidden(X3,X4))))))), file('data/MPTP/proofs/7d1ac91fca294d128bb8aed501eb880b.E_input', t45_pre_topc)).
fof(t49_pre_topc, conjecture, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r1_tarski(X2,X3)=>r1_tarski(k2_pre_topc(X1,X2),k2_pre_topc(X1,X3)))))), file('data/MPTP/proofs/7d1ac91fca294d128bb8aed501eb880b.E_input', t49_pre_topc)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/7d1ac91fca294d128bb8aed501eb880b.E_input', t3_subset)).
fof(dt_k2_pre_topc, axiom, ![X1, X2]:((l1_pre_topc(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k2_pre_topc(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/7d1ac91fca294d128bb8aed501eb880b.E_input', dt_k2_pre_topc)).
fof(c_0_5, plain, ![X15, X16, X17, X18, X19]:((~r1_tarski(X15,X16)|(~r2_hidden(X17,X15)|r2_hidden(X17,X16)))&((r2_hidden(esk5_2(X18,X19),X18)|r1_tarski(X18,X19))&(~r2_hidden(esk5_2(X18,X19),X19)|r1_tarski(X18,X19)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_6, plain, ![X8, X9, X10, X11]:((~r2_hidden(X10,k2_pre_topc(X8,X9))|(~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X8)))|(~v4_pre_topc(X11,X8)|~r1_tarski(X9,X11)|r2_hidden(X10,X11)))|~r2_hidden(X10,u1_struct_0(X8))|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X8)))|~l1_pre_topc(X8))&((m1_subset_1(esk4_3(X8,X9,X10),k1_zfmisc_1(u1_struct_0(X8)))|r2_hidden(X10,k2_pre_topc(X8,X9))|~r2_hidden(X10,u1_struct_0(X8))|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X8)))|~l1_pre_topc(X8))&(((v4_pre_topc(esk4_3(X8,X9,X10),X8)|r2_hidden(X10,k2_pre_topc(X8,X9))|~r2_hidden(X10,u1_struct_0(X8))|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X8)))|~l1_pre_topc(X8))&(r1_tarski(X9,esk4_3(X8,X9,X10))|r2_hidden(X10,k2_pre_topc(X8,X9))|~r2_hidden(X10,u1_struct_0(X8))|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X8)))|~l1_pre_topc(X8)))&(~r2_hidden(X10,esk4_3(X8,X9,X10))|r2_hidden(X10,k2_pre_topc(X8,X9))|~r2_hidden(X10,u1_struct_0(X8))|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X8)))|~l1_pre_topc(X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t45_pre_topc])])])])])).
fof(c_0_7, negated_conjecture, ~(![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r1_tarski(X2,X3)=>r1_tarski(k2_pre_topc(X1,X2),k2_pre_topc(X1,X3))))))), inference(assume_negation,[status(cth)],[t49_pre_topc])).
cnf(c_0_8, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (r1_tarski(X1,esk4_3(X2,X1,X3))|r2_hidden(X3,k2_pre_topc(X2,X1))|~r2_hidden(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_10, plain, ![X21, X22]:((~m1_subset_1(X21,k1_zfmisc_1(X22))|r1_tarski(X21,X22))&(~r1_tarski(X21,X22)|m1_subset_1(X21,k1_zfmisc_1(X22)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_11, negated_conjecture, (l1_pre_topc(esk1_0)&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(r1_tarski(esk2_0,esk3_0)&~r1_tarski(k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
cnf(c_0_12, plain, (r2_hidden(X1,k2_pre_topc(X2,X3))|~r2_hidden(X1,esk4_3(X2,X3,X1))|~r2_hidden(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_13, plain, (r2_hidden(X1,esk4_3(X2,X3,X4))|r2_hidden(X4,k2_pre_topc(X2,X3))|~r2_hidden(X4,u1_struct_0(X2))|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_14, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_17, plain, (r2_hidden(X1,k2_pre_topc(X2,X3))|~r2_hidden(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_18, negated_conjecture, (r1_tarski(esk3_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_19, plain, (r1_tarski(X1,k2_pre_topc(X2,X3))|~r2_hidden(esk5_2(X1,k2_pre_topc(X2,X3)),u1_struct_0(X2))|~r2_hidden(esk5_2(X1,k2_pre_topc(X2,X3)),X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_20, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_8, c_0_18])).
cnf(c_0_21, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_22, negated_conjecture, (r1_tarski(X1,k2_pre_topc(esk1_0,X2))|~r2_hidden(esk5_2(X1,k2_pre_topc(esk1_0,X2)),esk3_0)|~r2_hidden(esk5_2(X1,k2_pre_topc(esk1_0,X2)),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])])).
cnf(c_0_23, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_24, negated_conjecture, (r1_tarski(esk3_0,k2_pre_topc(esk1_0,X1))|~r2_hidden(esk5_2(esk3_0,k2_pre_topc(esk1_0,X1)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_25, negated_conjecture, (r1_tarski(esk3_0,k2_pre_topc(esk1_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_23]), c_0_15])])).
cnf(c_0_26, negated_conjecture, (r2_hidden(X1,k2_pre_topc(esk1_0,esk3_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_8, c_0_25])).
cnf(c_0_27, negated_conjecture, (r1_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_28, negated_conjecture, (r1_tarski(X1,k2_pre_topc(esk1_0,esk3_0))|~r2_hidden(esk5_2(X1,k2_pre_topc(esk1_0,esk3_0)),esk3_0)), inference(spm,[status(thm)],[c_0_16, c_0_26])).
cnf(c_0_29, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_8, c_0_27])).
cnf(c_0_30, negated_conjecture, (r1_tarski(X1,k2_pre_topc(esk1_0,esk3_0))|~r2_hidden(esk5_2(X1,k2_pre_topc(esk1_0,esk3_0)),esk2_0)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_32, plain, ![X13, X14]:(~l1_pre_topc(X13)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|m1_subset_1(k2_pre_topc(X13,X14),k1_zfmisc_1(u1_struct_0(X13)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_pre_topc])])).
cnf(c_0_33, negated_conjecture, (r1_tarski(esk2_0,k2_pre_topc(esk1_0,esk3_0))), inference(spm,[status(thm)],[c_0_30, c_0_23])).
cnf(c_0_34, negated_conjecture, (r1_tarski(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_14, c_0_31])).
cnf(c_0_35, plain, (m1_subset_1(k2_pre_topc(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_36, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,k2_pre_topc(X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v4_pre_topc(X4,X2)|~r1_tarski(X3,X4)|~r2_hidden(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_37, negated_conjecture, (r2_hidden(X1,k2_pre_topc(esk1_0,esk3_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_8, c_0_33])).
cnf(c_0_38, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_8, c_0_34])).
cnf(c_0_39, plain, (r1_tarski(k2_pre_topc(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_14, c_0_35])).
cnf(c_0_40, negated_conjecture, (r2_hidden(X1,X2)|~v4_pre_topc(X2,esk1_0)|~r2_hidden(X1,esk2_0)|~r1_tarski(esk3_0,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_15]), c_0_21])]), c_0_38])).
cnf(c_0_41, plain, (v4_pre_topc(esk4_3(X1,X2,X3),X1)|r2_hidden(X3,k2_pre_topc(X1,X2))|~r2_hidden(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_42, plain, (r2_hidden(X1,u1_struct_0(X2))|~r2_hidden(X1,k2_pre_topc(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_8, c_0_39])).
cnf(c_0_43, negated_conjecture, (r2_hidden(X1,esk4_3(esk1_0,X2,X3))|r2_hidden(X3,k2_pre_topc(esk1_0,X2))|~r2_hidden(X3,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)|~r1_tarski(esk3_0,esk4_3(esk1_0,X2,X3))|~m1_subset_1(esk4_3(esk1_0,X2,X3),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_21])])).
cnf(c_0_44, plain, (m1_subset_1(esk4_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|r2_hidden(X3,k2_pre_topc(X1,X2))|~r2_hidden(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_45, plain, (r2_hidden(esk5_2(k2_pre_topc(X1,X2),X3),u1_struct_0(X1))|r1_tarski(k2_pre_topc(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_42, c_0_23])).
cnf(c_0_46, negated_conjecture, (r2_hidden(X1,esk4_3(esk1_0,X2,X3))|r2_hidden(X3,k2_pre_topc(esk1_0,X2))|~r2_hidden(X3,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)|~r1_tarski(esk3_0,esk4_3(esk1_0,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_21])])).
cnf(c_0_47, plain, (r2_hidden(esk5_2(k2_pre_topc(X1,X2),X3),X4)|r1_tarski(k2_pre_topc(X1,X2),X3)|~v4_pre_topc(X4,X1)|~r1_tarski(X2,X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_23]), c_0_45])).
cnf(c_0_48, negated_conjecture, (r2_hidden(X1,esk4_3(esk1_0,esk3_0,X2))|r2_hidden(X2,k2_pre_topc(esk1_0,esk3_0))|~r2_hidden(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_9]), c_0_15]), c_0_21])])).
cnf(c_0_49, plain, (r2_hidden(esk5_2(k2_pre_topc(X1,X2),X3),k2_pre_topc(X4,X5))|r1_tarski(k2_pre_topc(X1,X2),X3)|~v4_pre_topc(esk4_3(X4,X5,esk5_2(k2_pre_topc(X1,X2),X3)),X1)|~r2_hidden(esk5_2(k2_pre_topc(X1,X2),X3),u1_struct_0(X4))|~r1_tarski(X2,esk4_3(X4,X5,esk5_2(k2_pre_topc(X1,X2),X3)))|~m1_subset_1(esk4_3(X4,X5,esk5_2(k2_pre_topc(X1,X2),X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X4)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_12, c_0_47])).
cnf(c_0_50, negated_conjecture, (r2_hidden(X1,k2_pre_topc(esk1_0,esk3_0))|r1_tarski(X2,esk4_3(esk1_0,esk3_0,X1))|~r2_hidden(esk5_2(X2,esk4_3(esk1_0,esk3_0,X1)),esk2_0)|~r2_hidden(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_16, c_0_48])).
cnf(c_0_51, plain, (r2_hidden(esk5_2(k2_pre_topc(X1,X2),X3),k2_pre_topc(X1,X4))|r1_tarski(k2_pre_topc(X1,X2),X3)|~r1_tarski(X2,esk4_3(X1,X4,esk5_2(k2_pre_topc(X1,X2),X3)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_41]), c_0_44]), c_0_45])).
cnf(c_0_52, negated_conjecture, (r2_hidden(X1,k2_pre_topc(esk1_0,esk3_0))|r1_tarski(esk2_0,esk4_3(esk1_0,esk3_0,X1))|~r2_hidden(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_50, c_0_23])).
cnf(c_0_53, negated_conjecture, (r2_hidden(esk5_2(k2_pre_topc(esk1_0,esk2_0),X1),k2_pre_topc(esk1_0,esk3_0))|r1_tarski(k2_pre_topc(esk1_0,esk2_0),X1)|~r2_hidden(esk5_2(k2_pre_topc(esk1_0,esk2_0),X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_15]), c_0_31]), c_0_21])])).
cnf(c_0_54, negated_conjecture, (~r1_tarski(k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_55, negated_conjecture, (~r2_hidden(esk5_2(k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_53]), c_0_54])).
cnf(c_0_56, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_45]), c_0_31]), c_0_21])]), c_0_54]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 57
# Proof object clause steps            : 46
# Proof object formula steps           : 11
# Proof object conjectures             : 30
# Proof object clause conjectures      : 27
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 31
# Proof object simplifying inferences  : 27
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 16
# Processed clauses                    : 662
# ...of these trivial                  : 1
# ...subsumed                          : 319
# ...remaining for further processing  : 342
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 64
# Backward-rewritten                   : 0
# Generated clauses                    : 1598
# ...of the previous two non-trivial   : 1571
# Contextual simplify-reflections      : 44
# Paramodulations                      : 1598
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 262
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 245
# Current number of unprocessed clauses: 916
# ...number of literals in the above   : 7281
# Current number of archived formulas  : 0
# Current number of archived clauses   : 80
# Clause-clause subsumption calls (NU) : 18903
# Rec. Clause-clause subsumption calls : 5918
# Non-unit clause-clause subsumptions  : 342
# Unit Clause-clause subsumption calls : 257
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 31
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 60113

# -------------------------------------------------
# User time                : 0.050 s
# System time              : 0.000 s
# Total time               : 0.050 s
# Maximum resident set size: 28516 pages
